fof(f5456, plain, $false, inference(avatar_sat_refutation, [], [f605, f622, f639, f656, f673, f690, f707, f724, f741, f758, f775, f792, f809, f826, f843, f860, f861, f862, f863, f864, f865, f866, f867, f868, f869, f870, f871, f872, f873, f874, f875, f876, f877, f878, f879, f880, f882, f883, f884, f885, f886, f887, f888, f889, f890, f891, f892, f897, f898, f900, f905, f906, f908, f913, f914, f916, f921, f922, f924, f929, f930, f931, f932, f937, f938, f939, f940, f945, f946, f947, f948, f953, f956, f961, f962, f963, f964, f969, f970, f972, f977, f978, f979, f980, f985, f986, f987, f988, f993, f994, f995, f996, f1001, f1002, f1003, f1004, f1009, f1010, f1012, f1017, f1018, f1019, f1020, f1025, f1028, f1033, f1034, f1035, f1036, f1041, f1042, f1043, f1044, f1049, f1050, f1052, f1057, f1060, f1065, f1066, f1068, f1073, f1074, f1076, f1081, f1082, f1083, f1084, f1089, f1092, f1097, f1098, f1105, f1106, f1107, f1108, f1113, f1114, f1115, f1116, f1121, f1122, f1124, f1129, f1130, f1131, f1132, f1137, f1138, f1140, f1145, f1146, f1153, f1154, f1155, f1156, f1161, f1162, f1164, f1169, f1170, f1171, f1172, f1177, f1178, f1179, f1180, f1185, f1186, f1187, f1193, f1194, f1196, f1201, f1202, f1203, f1204, f1209, f1210, f1212, f1217, f1218, f1220, f1225, f1228, f1233, f1234, f1236, f1241, f1242, f1243, f1244, f1249, f1250, f1251, f1252, f1257, f1258, f1260, f1265, f1266, f1273, f1274, f1276, f1281, f1282, f1283, f1284, f1289, f1290, f1291, f1292, f1297, f1298, f1300, f1305, f1306, f1307, f1308, f1313, f1314, f1315, f1316, f1321, f1322, f1323, f1324, f1329, f1332, f1337, f1338, f1339, f1340, f1345, f1346, f1347, f1348, f1353, f1354, f1356, f1361, f1364, f1369, f1370, f1372, f1377, f1378, f1380, f1385, f1386, f1388, f1393, f1396, f1401, f1402, f1404, f1409, f1410, f1411, f1412, f1417, f1418, f1419, f1420, f1421, f1424, f1441, f1458, f1475, f1492, f1493, f1495, f1496, f1505, f1514, f1519, f1528, f1533, f1538, f1547, f1556, f1557, f1566, f1567, f1572, f1577, f1586, f1587, f1596, f1598, f1603, f1609, f1614, f1615, f1616, f1630, f1633, f1655, f1666, f1676, f1678, f1702, f1725, f1740, f1752, f1768, f1771, f1774, f1780, f1783, f1802, f1818, f1832, f1841, f1847, f1853, f1863, f1866, f1871, f1874, f1898, f1900, f1933, f1963, f1970, f1973, f1999, f2007, f2014, f2045, f2051, f2063, f2067, f2070, f2093, f2099, f2102, f2107, f2109, f2118, f2123, f2141, f2147, f2162, f2169, f2175, f2177, f2222, f2227, f2228, f2255, f2256, f2265, f2271, f2282, f2289, f2293, f2295, f2298, f2311, f2324, f2339, f2345, f2346, f2373, f2388, f2404, f2413, f2415, f2419, f2421, f2445, f2449, f2455, f2469, f2508, f2521, f2523, f2536, f2546, f2562, f2564, f2601, f2612, f2627, f2631, f2638, f2657, f2702, f2714, f2718, f2723, f2751, f2765, f2774, f2787, f2821, f2831, f2849, f2876, f2882, f2887, f2905, f2907, f2909, f2927, f2942, f2947, f2979, f2998, f3008, f3009, f3037, f3060, f3069, f3105, f3136, f3149, f3150, f3189, f3212, f3216, f3225, f3242, f3258, f3279, f3281, f3284, f3294, f3305, f3351, f3356, f3362, f3371, f3381, f3397, f3414, f3430, f3435, f3440, f3442, f3470, f3482, f3495, f3502, f3541, f3563, f3567, f3577, f3580, f3603, f3611, f3642, f3649, f3672, f3676, f3679, f3707, f3708, f3727, f3735, f3752, f3764, f3789, f3791, f3801, f3829, f3870, f3873, f3875, f3893, f3894, f3911, f3921, f3924, f3938, f3940, f3958, f3966, f3981, f3983, f3995, f3996, f3997, f4014, f4015, f4019, f4038, f4041, f4043, f4048, f4051, f4055, f4075, f4109, f4110, f4116, f4120, f4130, f4139, f4143, f4146, f4163, f4167, f4170, f4172, f4176, f4196, f4197, f4216, f4217, f4218, f4227, f4240, f4252, f4278, f4282, f4295, f4304, f4313, f4319, f4331, f4338, f4339, f4347, f4352, f4381, f4383, f4386, f4390, f4415, f4417, f4428, f4439, f4447, f4454, f4464, f4472, f4480, f4493, f4501, f4505, f4545, f4550, f4552, f4564, f4565, f4571, f4574, f4582, f4583, f4589, f4593, f4602, f4620, f4622, f4625, f4637, f4638, f4649, f4655, f4672, f4683, f4696, f4704, f4710, f4716, f4722, f4727, f4731, f4745, f4748, f4760, f4763, f4787, f4793, f4808, f4809, f4816, f4821, f4822, f4838, f4844, f4849, f4857, f4865, f4874, f4875, f4883, f4890, f4891, f4903, f4941, f4955, f4971, f4984, f4985, f4987, f4990, f4998, f4999, f5007, f5011, f5027, f5031, f5045, f5061, f5062, f5064, f5070, f5071, f5080, f5082, f5087, f5096, f5101, f5107, f5116, f5122, f5133, f5142, f5144, f5170, f5171, f5172, f5175, f5183, f5184, f5198, f5207, f5210, f5217, f5244, f5245, f5249, f5256, f5265, f5266, f5275, f5276, f5281, f5285, f5291, f5297, f5302, f5309, f5326, f5338, f5339, f5346, f5347, f5353, f5358, f5365, f5376, f5380, f5385, f5388, f5393, f5404, f5432, f5433, f5438, f5442, f5455])).
fof(f5455, plain, (~ spl66_17 | ~ spl66_24 | spl66_159), inference(avatar_contradiction_clause, [], [f5454])).
fof(f5454, plain, ($false | (~ spl66_17 | ~ spl66_24 | spl66_159)), inference(subsumption_resolution, [], [f5448, f660])).
fof(f660, plain, ((e0 = op(e2, e3)) | ~ spl66_17), inference(avatar_component_clause, [], [f658])).
fof(f658, plain, (spl66_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_17])])).
fof(f5448, plain, (~ (e0 = op(e2, e3)) | (~ spl66_24 | spl66_159)), inference(backward_demodulation, [], [f1555, f689])).
fof(f689, plain, ((e3 = op(e2, e2)) | ~ spl66_24), inference(avatar_component_clause, [], [f687])).
fof(f687, plain, (spl66_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_24])])).
fof(f1555, plain, (~ (e0 = op(e2, op(e2, e2))) | spl66_159), inference(avatar_component_clause, [], [f1553])).
fof(f1553, plain, (spl66_159 <=> (e0 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_159])])).
fof(f5442, plain, (~ spl66_23 | ~ spl66_31), inference(avatar_split_clause, [], [f5439, f717, f683])).
fof(f683, plain, (spl66_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_23])])).
fof(f717, plain, (spl66_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_31])])).
fof(f5439, plain, (~ (e2 = op(e2, e2)) | ~ spl66_31), inference(backward_demodulation, [], [f272, f719])).
fof(f719, plain, ((e2 = op(e2, e0)) | ~ spl66_31), inference(avatar_component_clause, [], [f717])).
fof(f272, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax3)).
fof(f5438, plain, (~ spl66_23 | ~ spl66_39), inference(avatar_split_clause, [], [f5435, f751, f683])).
fof(f751, plain, (spl66_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_39])])).
fof(f5435, plain, (~ (e2 = op(e2, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f250, f753])).
fof(f753, plain, ((e2 = op(e1, e2)) | ~ spl66_39), inference(avatar_component_clause, [], [f751])).
fof(f250, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f5433, plain, (~ spl66_40 | ~ spl66_48), inference(avatar_split_clause, [], [f5429, f789, f755])).
fof(f755, plain, (spl66_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_40])])).
fof(f789, plain, (spl66_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_48])])).
fof(f5429, plain, (~ (e3 = op(e1, e2)) | ~ spl66_48), inference(backward_demodulation, [], [f266, f791])).
fof(f791, plain, ((e3 = op(e1, e0)) | ~ spl66_48), inference(avatar_component_clause, [], [f789])).
fof(f266, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f5432, plain, (~ spl66_32 | ~ spl66_48), inference(avatar_split_clause, [], [f5428, f789, f721])).
fof(f721, plain, (spl66_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_32])])).
fof(f5428, plain, (~ (e3 = op(e2, e0)) | ~ spl66_48), inference(backward_demodulation, [], [f238, f791])).
fof(f238, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f5404, plain, (~ spl66_32 | ~ spl66_17 | spl66_138), inference(avatar_split_clause, [], [f5397, f1455, f658, f721])).
fof(f1455, plain, (spl66_138 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_138])])).
fof(f5397, plain, (~ (e3 = op(e2, e0)) | (~ spl66_17 | spl66_138)), inference(backward_demodulation, [], [f1456, f660])).
fof(f1456, plain, (~ (e3 = op(e2, op(e2, e3))) | spl66_138), inference(avatar_component_clause, [], [f1455])).
fof(f5393, plain, (~ spl66_18 | ~ spl66_26), inference(avatar_split_clause, [], [f5392, f696, f662])).
fof(f662, plain, (spl66_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_18])])).
fof(f696, plain, (spl66_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_26])])).
fof(f5392, plain, (~ (e1 = op(e2, e3)) | ~ spl66_26), inference(backward_demodulation, [], [f275, f698])).
fof(f698, plain, ((e1 = op(e2, e1)) | ~ spl66_26), inference(avatar_component_clause, [], [f696])).
fof(f275, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5388, plain, (~ spl66_34 | ~ spl66_41 | ~ spl66_142), inference(avatar_contradiction_clause, [], [f5387])).
fof(f5387, plain, ($false | (~ spl66_34 | ~ spl66_41 | ~ spl66_142)), inference(subsumption_resolution, [], [f5386, f285])).
fof(f285, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax4)).
fof(f5386, plain, ((e0 = e3) | (~ spl66_34 | ~ spl66_41 | ~ spl66_142)), inference(forward_demodulation, [], [f5384, f762])).
fof(f762, plain, ((e0 = op(e1, e1)) | ~ spl66_41), inference(avatar_component_clause, [], [f760])).
fof(f760, plain, (spl66_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_41])])).
fof(f5384, plain, ((e3 = op(e1, e1)) | (~ spl66_34 | ~ spl66_142)), inference(backward_demodulation, [], [f1474, f732])).
fof(f732, plain, ((e1 = op(e1, e3)) | ~ spl66_34), inference(avatar_component_clause, [], [f730])).
fof(f730, plain, (spl66_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_34])])).
fof(f1474, plain, ((e3 = op(e1, op(e1, e3))) | ~ spl66_142), inference(avatar_component_clause, [], [f1472])).
fof(f1472, plain, (spl66_142 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_142])])).
fof(f5385, plain, (~ spl66_18 | ~ spl66_34), inference(avatar_split_clause, [], [f5381, f730, f662])).
fof(f5381, plain, (~ (e1 = op(e2, e3)) | ~ spl66_34), inference(backward_demodulation, [], [f256, f732])).
fof(f256, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5380, plain, (~ spl66_35 | ~ spl66_40 | spl66_141), inference(avatar_split_clause, [], [f5379, f1468, f755, f734])).
fof(f734, plain, (spl66_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_35])])).
fof(f1468, plain, (spl66_141 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_141])])).
fof(f5379, plain, (~ (e2 = op(e1, e3)) | (~ spl66_40 | spl66_141)), inference(backward_demodulation, [], [f1469, f757])).
fof(f757, plain, ((e3 = op(e1, e2)) | ~ spl66_40), inference(avatar_component_clause, [], [f755])).
fof(f1469, plain, (~ (e2 = op(e1, op(e1, e2))) | spl66_141), inference(avatar_component_clause, [], [f1468])).
fof(f5376, plain, (~ spl66_25 | ~ spl66_41), inference(avatar_split_clause, [], [f5366, f760, f692])).
fof(f692, plain, (spl66_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_25])])).
fof(f5366, plain, (~ (e0 = op(e2, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f244, f762])).
fof(f244, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f5365, plain, (~ spl66_35 | ~ spl66_47), inference(avatar_split_clause, [], [f5362, f785, f734])).
fof(f785, plain, (spl66_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_47])])).
fof(f5362, plain, (~ (e2 = op(e1, e3)) | ~ spl66_47), inference(backward_demodulation, [], [f267, f787])).
fof(f787, plain, ((e2 = op(e1, e0)) | ~ spl66_47), inference(avatar_component_clause, [], [f785])).
fof(f267, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f5358, plain, (~ spl66_35 | ~ spl66_51), inference(avatar_split_clause, [], [f5354, f802, f734])).
fof(f802, plain, (spl66_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_51])])).
fof(f5354, plain, (~ (e2 = op(e1, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f253, f804])).
fof(f804, plain, ((e2 = op(e0, e3)) | ~ spl66_51), inference(avatar_component_clause, [], [f802])).
fof(f253, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f5353, plain, (~ spl66_38 | ~ spl66_54), inference(avatar_split_clause, [], [f5348, f815, f747])).
fof(f747, plain, (spl66_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_38])])).
fof(f815, plain, (spl66_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_54])])).
fof(f5348, plain, (~ (e1 = op(e1, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f247, f817])).
fof(f817, plain, ((e1 = op(e0, e2)) | ~ spl66_54), inference(avatar_component_clause, [], [f815])).
fof(f247, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f5347, plain, (~ spl66_56 | ~ spl66_60), inference(avatar_split_clause, [], [f5343, f840, f823])).
fof(f823, plain, (spl66_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_56])])).
fof(f840, plain, (spl66_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_60])])).
fof(f5343, plain, (~ (e3 = op(e0, e2)) | ~ spl66_60), inference(backward_demodulation, [], [f262, f842])).
fof(f842, plain, ((e3 = op(e0, e1)) | ~ spl66_60), inference(avatar_component_clause, [], [f840])).
fof(f262, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f5346, plain, (~ spl66_44 | ~ spl66_60), inference(avatar_split_clause, [], [f5341, f840, f772])).
fof(f772, plain, (spl66_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_44])])).
fof(f5341, plain, (~ (e3 = op(e1, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f241, f842])).
fof(f241, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f5339, plain, (~ spl66_49 | ~ spl66_61), inference(avatar_split_clause, [], [f5330, f845, f794])).
fof(f794, plain, (spl66_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_49])])).
fof(f845, plain, (spl66_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_61])])).
fof(f5330, plain, (~ (e0 = op(e0, e3)) | ~ spl66_61), inference(backward_demodulation, [], [f261, f847])).
fof(f847, plain, ((e0 = op(e0, e0)) | ~ spl66_61), inference(avatar_component_clause, [], [f845])).
fof(f261, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f5338, plain, (~ spl66_45 | ~ spl66_61), inference(avatar_split_clause, [], [f5327, f845, f777])).
fof(f777, plain, (spl66_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_45])])).
fof(f5327, plain, (~ (e0 = op(e1, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f235, f847])).
fof(f235, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f5326, plain, (spl66_17 | ~ spl66_32 | ~ spl66_135), inference(avatar_split_clause, [], [f5325, f1443, f721, f658])).
fof(f1443, plain, (spl66_135 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_135])])).
fof(f5325, plain, ((e0 = op(e2, e3)) | (~ spl66_32 | ~ spl66_135)), inference(forward_demodulation, [], [f1445, f723])).
fof(f723, plain, ((e3 = op(e2, e0)) | ~ spl66_32), inference(avatar_component_clause, [], [f721])).
fof(f1445, plain, ((e0 = op(e2, op(e2, e0))) | ~ spl66_135), inference(avatar_component_clause, [], [f1443])).
fof(f5309, plain, (~ spl66_20 | ~ spl66_32), inference(avatar_split_clause, [], [f5308, f721, f670])).
fof(f670, plain, (spl66_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_20])])).
fof(f5308, plain, (~ (e3 = op(e2, e3)) | ~ spl66_32), inference(forward_demodulation, [], [f273, f723])).
fof(f273, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5302, plain, (~ spl66_6 | ~ spl66_11 | spl66_132), inference(avatar_split_clause, [], [f5301, f1430, f632, f611])).
fof(f611, plain, (spl66_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_6])])).
fof(f632, plain, (spl66_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_11])])).
fof(f1430, plain, (spl66_132 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_132])])).
fof(f5301, plain, (~ (e1 = op(e3, e2)) | (~ spl66_11 | spl66_132)), inference(backward_demodulation, [], [f1431, f634])).
fof(f634, plain, ((e2 = op(e3, e1)) | ~ spl66_11), inference(avatar_component_clause, [], [f632])).
fof(f1431, plain, (~ (e1 = op(e3, op(e3, e1))) | spl66_132), inference(avatar_component_clause, [], [f1430])).
fof(f5297, plain, (~ spl66_6 | ~ spl66_14), inference(avatar_split_clause, [], [f5293, f645, f611])).
fof(f645, plain, (spl66_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_14])])).
fof(f5293, plain, (~ (e1 = op(e3, e2)) | ~ spl66_14), inference(backward_demodulation, [], [f278, f647])).
fof(f647, plain, ((e1 = op(e3, e0)) | ~ spl66_14), inference(avatar_component_clause, [], [f645])).
fof(f278, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5291, plain, (~ spl66_9 | ~ spl66_25), inference(avatar_split_clause, [], [f5286, f692, f624])).
fof(f624, plain, (spl66_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_9])])).
fof(f5286, plain, (~ (e0 = op(e3, e1)) | ~ spl66_25), inference(backward_demodulation, [], [f246, f694])).
fof(f694, plain, ((e0 = op(e2, e1)) | ~ spl66_25), inference(avatar_component_clause, [], [f692])).
fof(f246, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f5285, plain, (~ spl66_28 | ~ spl66_32), inference(avatar_split_clause, [], [f5283, f721, f704])).
fof(f704, plain, (spl66_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_28])])).
fof(f5283, plain, (~ (e3 = op(e2, e1)) | ~ spl66_32), inference(backward_demodulation, [], [f271, f723])).
fof(f271, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f5281, plain, (~ spl66_6 | ~ spl66_38), inference(avatar_split_clause, [], [f5280, f747, f611])).
fof(f5280, plain, (~ (e1 = op(e3, e2)) | ~ spl66_38), inference(backward_demodulation, [], [f251, f749])).
fof(f749, plain, ((e1 = op(e1, e2)) | ~ spl66_38), inference(avatar_component_clause, [], [f747])).
fof(f251, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5276, plain, (~ spl66_5 | ~ spl66_44 | spl66_152 | ~ spl66_153), inference(avatar_split_clause, [], [f5274, f1525, f1521, f772, f607])).
fof(f607, plain, (spl66_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_5])])).
fof(f1521, plain, (spl66_152 <=> (e0 = op(op(e1, e1), op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl66_152])])).
fof(f1525, plain, (spl66_153 <=> (e2 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_153])])).
fof(f5274, plain, (~ (e0 = op(e3, e2)) | (~ spl66_44 | spl66_152 | ~ spl66_153)), inference(backward_demodulation, [], [f5228, f774])).
fof(f774, plain, ((e3 = op(e1, e1)) | ~ spl66_44), inference(avatar_component_clause, [], [f772])).
fof(f5228, plain, (~ (e0 = op(op(e1, e1), e2)) | (spl66_152 | ~ spl66_153)), inference(forward_demodulation, [], [f1523, f1526])).
fof(f1526, plain, ((e2 = op(e1, op(e1, e1))) | ~ spl66_153), inference(avatar_component_clause, [], [f1525])).
fof(f1523, plain, (~ (e0 = op(op(e1, e1), op(e1, op(e1, e1)))) | spl66_152), inference(avatar_component_clause, [], [f1521])).
fof(f5275, plain, (~ spl66_28 | ~ spl66_44), inference(avatar_split_clause, [], [f5267, f772, f704])).
fof(f5267, plain, (~ (e3 = op(e2, e1)) | ~ spl66_44), inference(backward_demodulation, [], [f244, f774])).
fof(f5266, plain, (~ spl66_37 | ~ spl66_45), inference(avatar_split_clause, [], [f5260, f777, f743])).
fof(f743, plain, (spl66_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_37])])).
fof(f5260, plain, (~ (e0 = op(e1, e2)) | ~ spl66_45), inference(backward_demodulation, [], [f266, f779])).
fof(f779, plain, ((e0 = op(e1, e0)) | ~ spl66_45), inference(avatar_component_clause, [], [f777])).
fof(f5265, plain, (~ spl66_29 | ~ spl66_45), inference(avatar_split_clause, [], [f5257, f777, f709])).
fof(f709, plain, (spl66_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_29])])).
fof(f5257, plain, (~ (e0 = op(e2, e0)) | ~ spl66_45), inference(backward_demodulation, [], [f238, f779])).
fof(f5256, plain, (~ spl66_42 | ~ spl66_58), inference(avatar_split_clause, [], [f5250, f832, f764])).
fof(f764, plain, (spl66_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_42])])).
fof(f832, plain, (spl66_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_58])])).
fof(f5250, plain, (~ (e1 = op(e1, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f241, f834])).
fof(f834, plain, ((e1 = op(e0, e1)) | ~ spl66_58), inference(avatar_component_clause, [], [f832])).
fof(f5249, plain, (~ spl66_18 | ~ spl66_63 | spl66_160 | ~ spl66_162), inference(avatar_contradiction_clause, [], [f5248])).
fof(f5248, plain, ($false | (~ spl66_18 | ~ spl66_63 | spl66_160 | ~ spl66_162)), inference(subsumption_resolution, [], [f5243, f664])).
fof(f664, plain, ((e1 = op(e2, e3)) | ~ spl66_18), inference(avatar_component_clause, [], [f662])).
fof(f5243, plain, (~ (e1 = op(e2, e3)) | (~ spl66_63 | spl66_160 | ~ spl66_162)), inference(backward_demodulation, [], [f5225, f855])).
fof(f855, plain, ((op(e0, e0) = e2) | ~ spl66_63), inference(avatar_component_clause, [], [f853])).
fof(f853, plain, (spl66_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl66_63])])).
fof(f5225, plain, (~ (e1 = op(op(e0, e0), e3)) | (spl66_160 | ~ spl66_162)), inference(forward_demodulation, [], [f1561, f1570])).
fof(f1570, plain, ((e3 = op(e0, op(e0, e0))) | ~ spl66_162), inference(avatar_component_clause, [], [f1569])).
fof(f1569, plain, (spl66_162 <=> (e3 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_162])])).
fof(f1561, plain, (~ (e1 = op(op(e0, e0), op(e0, op(e0, e0)))) | spl66_160), inference(avatar_component_clause, [], [f1559])).
fof(f1559, plain, (spl66_160 <=> (e1 = op(op(e0, e0), op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl66_160])])).
fof(f5245, plain, (~ spl66_59 | ~ spl66_63), inference(avatar_split_clause, [], [f5236, f853, f836])).
fof(f836, plain, (spl66_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_59])])).
fof(f5236, plain, (~ (e2 = op(e0, e1)) | ~ spl66_63), inference(backward_demodulation, [], [f259, f855])).
fof(f259, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f5244, plain, (~ spl66_15 | ~ spl66_63), inference(avatar_split_clause, [], [f5235, f853, f649])).
fof(f649, plain, (spl66_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_15])])).
fof(f5235, plain, (~ (e2 = op(e3, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f237, f855])).
fof(f237, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f5217, plain, (~ spl66_30 | ~ spl66_18), inference(avatar_split_clause, [], [f5216, f662, f713])).
fof(f713, plain, (spl66_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_30])])).
fof(f5216, plain, (~ (e1 = op(e2, e0)) | ~ spl66_18), inference(forward_demodulation, [], [f273, f664])).
fof(f5210, plain, (~ spl66_8 | ~ spl66_56), inference(avatar_split_clause, [], [f5209, f823, f619])).
fof(f619, plain, (spl66_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_8])])).
fof(f5209, plain, (~ (e3 = op(e3, e2)) | ~ spl66_56), inference(forward_demodulation, [], [f249, f825])).
fof(f825, plain, ((e3 = op(e0, e2)) | ~ spl66_56), inference(avatar_component_clause, [], [f823])).
fof(f249, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5207, plain, (~ spl66_49 | ~ spl66_56 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f5206])).
fof(f5206, plain, ($false | (~ spl66_49 | ~ spl66_56 | ~ spl66_145)), inference(subsumption_resolution, [], [f5205, f284])).
fof(f284, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f5205, plain, ((e0 = e2) | (~ spl66_49 | ~ spl66_56 | ~ spl66_145)), inference(forward_demodulation, [], [f5204, f796])).
fof(f796, plain, ((e0 = op(e0, e3)) | ~ spl66_49), inference(avatar_component_clause, [], [f794])).
fof(f5204, plain, ((e2 = op(e0, e3)) | (~ spl66_56 | ~ spl66_145)), inference(forward_demodulation, [], [f1487, f825])).
fof(f1487, plain, ((e2 = op(e0, op(e0, e2))) | ~ spl66_145), inference(avatar_component_clause, [], [f1485])).
fof(f1485, plain, (spl66_145 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_145])])).
fof(f5198, plain, (~ spl66_17 | ~ spl66_29), inference(avatar_split_clause, [], [f5197, f709, f658])).
fof(f5197, plain, (~ (e0 = op(e2, e3)) | ~ spl66_29), inference(backward_demodulation, [], [f273, f711])).
fof(f711, plain, ((e0 = op(e2, e0)) | ~ spl66_29), inference(avatar_component_clause, [], [f709])).
fof(f5184, plain, (~ spl66_17 | ~ spl66_49), inference(avatar_split_clause, [], [f5180, f794, f658])).
fof(f5180, plain, (~ (e0 = op(e2, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f254, f796])).
fof(f254, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5183, plain, (~ spl66_49 | ~ spl66_62 | ~ spl66_146), inference(avatar_contradiction_clause, [], [f5182])).
fof(f5182, plain, ($false | (~ spl66_49 | ~ spl66_62 | ~ spl66_146)), inference(subsumption_resolution, [], [f5181, f287])).
fof(f287, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f5181, plain, ((e1 = e3) | (~ spl66_49 | ~ spl66_62 | ~ spl66_146)), inference(forward_demodulation, [], [f5177, f851])).
fof(f851, plain, ((op(e0, e0) = e1) | ~ spl66_62), inference(avatar_component_clause, [], [f849])).
fof(f849, plain, (spl66_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl66_62])])).
fof(f5177, plain, ((op(e0, e0) = e3) | (~ spl66_49 | ~ spl66_146)), inference(backward_demodulation, [], [f1491, f796])).
fof(f1491, plain, ((e3 = op(e0, op(e0, e3))) | ~ spl66_146), inference(avatar_component_clause, [], [f1489])).
fof(f1489, plain, (spl66_146 <=> (e3 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_146])])).
fof(f5175, plain, (~ spl66_40 | ~ spl66_56), inference(avatar_split_clause, [], [f5173, f823, f755])).
fof(f5173, plain, (~ (e3 = op(e1, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f247, f825])).
fof(f5172, plain, (~ spl66_40 | ~ spl66_62 | ~ spl66_161 | spl66_170), inference(avatar_split_clause, [], [f5168, f1611, f1563, f849, f755])).
fof(f1563, plain, (spl66_161 <=> (e2 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_161])])).
fof(f1611, plain, (spl66_170 <=> (e3 = op(op(e0, e0), op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl66_170])])).
fof(f5168, plain, (~ (e3 = op(e1, e2)) | (~ spl66_62 | ~ spl66_161 | spl66_170)), inference(backward_demodulation, [], [f5156, f851])).
fof(f5156, plain, (~ (e3 = op(op(e0, e0), e2)) | (~ spl66_161 | spl66_170)), inference(backward_demodulation, [], [f1613, f1564])).
fof(f1564, plain, ((e2 = op(e0, op(e0, e0))) | ~ spl66_161), inference(avatar_component_clause, [], [f1563])).
fof(f1613, plain, (~ (e3 = op(op(e0, e0), op(e0, op(e0, e0)))) | spl66_170), inference(avatar_component_clause, [], [f1611])).
fof(f5171, plain, (~ spl66_50 | ~ spl66_62), inference(avatar_split_clause, [], [f5163, f849, f798])).
fof(f798, plain, (spl66_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_50])])).
fof(f5163, plain, (~ (e1 = op(e0, e3)) | ~ spl66_62), inference(backward_demodulation, [], [f261, f851])).
fof(f5170, plain, (~ spl66_30 | ~ spl66_62), inference(avatar_split_clause, [], [f5161, f849, f713])).
fof(f5161, plain, (~ (e1 = op(e2, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f236, f851])).
fof(f236, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f5144, plain, (~ spl66_20 | ~ spl66_4), inference(avatar_split_clause, [], [f5143, f602, f670])).
fof(f602, plain, (spl66_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_4])])).
fof(f5143, plain, (~ (e3 = op(e2, e3)) | ~ spl66_4), inference(forward_demodulation, [], [f258, f604])).
fof(f604, plain, ((e3 = op(e3, e3)) | ~ spl66_4), inference(avatar_component_clause, [], [f602])).
fof(f258, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5142, plain, (~ spl66_19 | ~ spl66_23), inference(avatar_split_clause, [], [f5141, f683, f666])).
fof(f666, plain, (spl66_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_19])])).
fof(f5141, plain, (~ (e2 = op(e2, e3)) | ~ spl66_23), inference(forward_demodulation, [], [f276, f685])).
fof(f685, plain, ((e2 = op(e2, e2)) | ~ spl66_23), inference(avatar_component_clause, [], [f683])).
fof(f276, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5133, plain, (~ spl66_2 | ~ spl66_6), inference(avatar_split_clause, [], [f5129, f611, f594])).
fof(f594, plain, (spl66_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_2])])).
fof(f5129, plain, (~ (e1 = op(e3, e3)) | ~ spl66_6), inference(backward_demodulation, [], [f282, f613])).
fof(f613, plain, ((e1 = op(e3, e2)) | ~ spl66_6), inference(avatar_component_clause, [], [f611])).
fof(f282, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5122, plain, (~ spl66_11 | ~ spl66_15), inference(avatar_split_clause, [], [f5117, f649, f632])).
fof(f5117, plain, (~ (e2 = op(e3, e1)) | ~ spl66_15), inference(backward_demodulation, [], [f277, f651])).
fof(f651, plain, ((e2 = op(e3, e0)) | ~ spl66_15), inference(avatar_component_clause, [], [f649])).
fof(f277, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f5116, plain, (~ spl66_22 | ~ spl66_23 | ~ spl66_137 | spl66_158), inference(avatar_split_clause, [], [f5113, f1549, f1451, f683, f679])).
fof(f679, plain, (spl66_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_22])])).
fof(f1451, plain, (spl66_137 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_137])])).
fof(f1549, plain, (spl66_158 <=> (e1 = op(op(e2, e2), op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl66_158])])).
fof(f5113, plain, (~ (e1 = op(e2, e2)) | (~ spl66_23 | ~ spl66_137 | spl66_158)), inference(backward_demodulation, [], [f5048, f685])).
fof(f5048, plain, (~ (e1 = op(op(e2, e2), e2)) | (~ spl66_137 | spl66_158)), inference(backward_demodulation, [], [f1551, f1453])).
fof(f1453, plain, ((e2 = op(e2, op(e2, e2))) | ~ spl66_137), inference(avatar_component_clause, [], [f1451])).
fof(f1551, plain, (~ (e1 = op(op(e2, e2), op(e2, op(e2, e2)))) | spl66_158), inference(avatar_component_clause, [], [f1549])).
fof(f5107, plain, (~ spl66_22 | ~ spl66_30), inference(avatar_split_clause, [], [f5103, f713, f679])).
fof(f5103, plain, (~ (e1 = op(e2, e2)) | ~ spl66_30), inference(backward_demodulation, [], [f272, f715])).
fof(f715, plain, ((e1 = op(e2, e0)) | ~ spl66_30), inference(avatar_component_clause, [], [f713])).
fof(f5101, plain, (~ spl66_8 | ~ spl66_40), inference(avatar_split_clause, [], [f5099, f755, f619])).
fof(f5099, plain, (~ (e3 = op(e3, e2)) | ~ spl66_40), inference(backward_demodulation, [], [f251, f757])).
fof(f5096, plain, (~ spl66_13 | ~ spl66_45), inference(avatar_split_clause, [], [f5090, f777, f641])).
fof(f641, plain, (spl66_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_13])])).
fof(f5090, plain, (~ (e0 = op(e3, e0)) | ~ spl66_45), inference(backward_demodulation, [], [f239, f779])).
fof(f239, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f5087, plain, (~ spl66_2 | ~ spl66_50), inference(avatar_split_clause, [], [f5083, f798, f594])).
fof(f5083, plain, (~ (e1 = op(e3, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f255, f800])).
fof(f800, plain, ((e1 = op(e0, e3)) | ~ spl66_50), inference(avatar_component_clause, [], [f798])).
fof(f255, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5082, plain, (~ spl66_53 | spl66_63 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f5081])).
fof(f5081, plain, ($false | (~ spl66_53 | spl66_63 | ~ spl66_145)), inference(subsumption_resolution, [], [f5079, f854])).
fof(f854, plain, (~ (op(e0, e0) = e2) | spl66_63), inference(avatar_component_clause, [], [f853])).
fof(f5079, plain, ((op(e0, e0) = e2) | (~ spl66_53 | ~ spl66_145)), inference(backward_demodulation, [], [f1487, f813])).
fof(f813, plain, ((e0 = op(e0, e2)) | ~ spl66_53), inference(avatar_component_clause, [], [f811])).
fof(f811, plain, (spl66_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_53])])).
fof(f5080, plain, (~ spl66_37 | ~ spl66_53), inference(avatar_split_clause, [], [f5072, f811, f743])).
fof(f5072, plain, (~ (e0 = op(e1, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f247, f813])).
fof(f5071, plain, (~ spl66_55 | ~ spl66_59), inference(avatar_split_clause, [], [f5066, f836, f819])).
fof(f819, plain, (spl66_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_55])])).
fof(f5066, plain, (~ (e2 = op(e0, e2)) | ~ spl66_59), inference(backward_demodulation, [], [f262, f838])).
fof(f838, plain, ((e2 = op(e0, e1)) | ~ spl66_59), inference(avatar_component_clause, [], [f836])).
fof(f5070, plain, (~ spl66_11 | ~ spl66_59), inference(avatar_split_clause, [], [f5065, f836, f632])).
fof(f5065, plain, (~ (e2 = op(e3, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f243, f838])).
fof(f243, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f5064, plain, (~ spl66_11 | ~ spl66_64 | spl66_166 | ~ spl66_167), inference(avatar_split_clause, [], [f5060, f1593, f1589, f857, f632])).
fof(f857, plain, (spl66_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl66_64])])).
fof(f1589, plain, (spl66_166 <=> (e2 = op(op(e0, e0), op(e0, op(e0, e0))))), introduced(avatar_definition, [new_symbols(naming, [spl66_166])])).
fof(f1593, plain, (spl66_167 <=> (e1 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_167])])).
fof(f5060, plain, (~ (e2 = op(e3, e1)) | (~ spl66_64 | spl66_166 | ~ spl66_167)), inference(backward_demodulation, [], [f5043, f859])).
fof(f859, plain, ((op(e0, e0) = e3) | ~ spl66_64), inference(avatar_component_clause, [], [f857])).
fof(f5043, plain, (~ (e2 = op(op(e0, e0), e1)) | (spl66_166 | ~ spl66_167)), inference(backward_demodulation, [], [f1591, f1594])).
fof(f1594, plain, ((e1 = op(e0, op(e0, e0))) | ~ spl66_167), inference(avatar_component_clause, [], [f1593])).
fof(f1591, plain, (~ (e2 = op(op(e0, e0), op(e0, op(e0, e0)))) | spl66_166), inference(avatar_component_clause, [], [f1589])).
fof(f5062, plain, (~ spl66_52 | ~ spl66_64), inference(avatar_split_clause, [], [f5056, f857, f806])).
fof(f806, plain, (spl66_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_52])])).
fof(f5056, plain, (~ (e3 = op(e0, e3)) | ~ spl66_64), inference(backward_demodulation, [], [f261, f859])).
fof(f5061, plain, (~ spl66_48 | ~ spl66_64), inference(avatar_split_clause, [], [f5051, f857, f789])).
fof(f5051, plain, (~ (e3 = op(e1, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f235, f859])).
fof(f5045, plain, (spl66_40 | ~ spl66_35 | ~ spl66_142), inference(avatar_split_clause, [], [f5044, f1472, f734, f755])).
fof(f5044, plain, ((e3 = op(e1, e2)) | (~ spl66_35 | ~ spl66_142)), inference(forward_demodulation, [], [f1474, f736])).
fof(f736, plain, ((e2 = op(e1, e3)) | ~ spl66_35), inference(avatar_component_clause, [], [f734])).
fof(f5031, plain, (~ spl66_12 | ~ spl66_28), inference(avatar_split_clause, [], [f5030, f704, f636])).
fof(f636, plain, (spl66_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_12])])).
fof(f5030, plain, (~ (e3 = op(e3, e1)) | ~ spl66_28), inference(forward_demodulation, [], [f246, f706])).
fof(f706, plain, ((e3 = op(e2, e1)) | ~ spl66_28), inference(avatar_component_clause, [], [f704])).
fof(f5027, plain, (~ spl66_17 | ~ spl66_31 | ~ spl66_138), inference(avatar_contradiction_clause, [], [f5026])).
fof(f5026, plain, ($false | (~ spl66_17 | ~ spl66_31 | ~ spl66_138)), inference(subsumption_resolution, [], [f5025, f288])).
fof(f288, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f5025, plain, ((e2 = e3) | (~ spl66_17 | ~ spl66_31 | ~ spl66_138)), inference(forward_demodulation, [], [f5024, f719])).
fof(f5024, plain, ((e3 = op(e2, e0)) | (~ spl66_17 | ~ spl66_138)), inference(backward_demodulation, [], [f1457, f660])).
fof(f1457, plain, ((e3 = op(e2, op(e2, e3))) | ~ spl66_138), inference(avatar_component_clause, [], [f1455])).
fof(f5011, plain, (~ spl66_20 | ~ spl66_28), inference(avatar_split_clause, [], [f5009, f704, f670])).
fof(f5009, plain, (~ (e3 = op(e2, e3)) | ~ spl66_28), inference(backward_demodulation, [], [f275, f706])).
fof(f5007, plain, (~ spl66_21 | ~ spl66_37), inference(avatar_split_clause, [], [f5003, f743, f675])).
fof(f675, plain, (spl66_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_21])])).
fof(f5003, plain, (~ (e0 = op(e2, e2)) | ~ spl66_37), inference(backward_demodulation, [], [f250, f745])).
fof(f745, plain, ((e0 = op(e1, e2)) | ~ spl66_37), inference(avatar_component_clause, [], [f743])).
fof(f4999, plain, (~ spl66_38 | ~ spl66_42), inference(avatar_split_clause, [], [f4992, f764, f747])).
fof(f4992, plain, (~ (e1 = op(e1, e2)) | ~ spl66_42), inference(backward_demodulation, [], [f268, f766])).
fof(f766, plain, ((e1 = op(e1, e1)) | ~ spl66_42), inference(avatar_component_clause, [], [f764])).
fof(f268, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f4998, plain, (~ spl66_26 | ~ spl66_42), inference(avatar_split_clause, [], [f4991, f764, f696])).
fof(f4991, plain, (~ (e1 = op(e2, e1)) | ~ spl66_42), inference(backward_demodulation, [], [f244, f766])).
fof(f4990, plain, (~ spl66_20 | ~ spl66_52), inference(avatar_split_clause, [], [f4989, f806, f670])).
fof(f4989, plain, (~ (e3 = op(e2, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f254, f808])).
fof(f808, plain, ((e3 = op(e0, e3)) | ~ spl66_52), inference(avatar_component_clause, [], [f806])).
fof(f4987, plain, (~ spl66_57 | ~ spl66_62 | spl66_144), inference(avatar_contradiction_clause, [], [f4986])).
fof(f4986, plain, ($false | (~ spl66_57 | ~ spl66_62 | spl66_144)), inference(subsumption_resolution, [], [f4980, f851])).
fof(f4980, plain, (~ (op(e0, e0) = e1) | (~ spl66_57 | spl66_144)), inference(backward_demodulation, [], [f1482, f830])).
fof(f830, plain, ((e0 = op(e0, e1)) | ~ spl66_57), inference(avatar_component_clause, [], [f828])).
fof(f828, plain, (spl66_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_57])])).
fof(f1482, plain, (~ (e1 = op(e0, op(e0, e1))) | spl66_144), inference(avatar_component_clause, [], [f1481])).
fof(f1481, plain, (spl66_144 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_144])])).
fof(f4985, plain, (~ spl66_49 | ~ spl66_57), inference(avatar_split_clause, [], [f4976, f828, f794])).
fof(f4976, plain, (~ (e0 = op(e0, e3)) | ~ spl66_57), inference(backward_demodulation, [], [f263, f830])).
fof(f263, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4984, plain, (~ spl66_41 | ~ spl66_57), inference(avatar_split_clause, [], [f4974, f828, f760])).
fof(f4974, plain, (~ (e0 = op(e1, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f241, f830])).
fof(f4971, plain, (spl66_57 | ~ spl66_62 | ~ spl66_143), inference(avatar_split_clause, [], [f4970, f1477, f849, f828])).
fof(f1477, plain, (spl66_143 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_143])])).
fof(f4970, plain, ((e0 = op(e0, e1)) | (~ spl66_62 | ~ spl66_143)), inference(forward_demodulation, [], [f1479, f851])).
fof(f1479, plain, ((e0 = op(e0, op(e0, e0))) | ~ spl66_143), inference(avatar_component_clause, [], [f1477])).
fof(f4955, plain, (~ spl66_37 | ~ spl66_2 | spl66_147 | ~ spl66_151), inference(avatar_split_clause, [], [f4950, f1516, f1498, f594, f743])).
fof(f1498, plain, (spl66_147 <=> (e0 = op(op(e3, e3), op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl66_147])])).
fof(f1516, plain, (spl66_151 <=> (e2 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_151])])).
fof(f4950, plain, (~ (e0 = op(e1, e2)) | (~ spl66_2 | spl66_147 | ~ spl66_151)), inference(backward_demodulation, [], [f4860, f596])).
fof(f596, plain, ((e1 = op(e3, e3)) | ~ spl66_2), inference(avatar_component_clause, [], [f594])).
fof(f4860, plain, (~ (e0 = op(op(e3, e3), e2)) | (spl66_147 | ~ spl66_151)), inference(forward_demodulation, [], [f1500, f1517])).
fof(f1517, plain, ((e2 = op(e3, op(e3, e3))) | ~ spl66_151), inference(avatar_component_clause, [], [f1516])).
fof(f1500, plain, (~ (e0 = op(op(e3, e3), op(e3, op(e3, e3)))) | spl66_147), inference(avatar_component_clause, [], [f1498])).
fof(f4941, plain, (~ spl66_1 | ~ spl66_13), inference(avatar_split_clause, [], [f4936, f641, f590])).
fof(f590, plain, (spl66_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_1])])).
fof(f4936, plain, (~ (e0 = op(e3, e3)) | ~ spl66_13), inference(backward_demodulation, [], [f279, f643])).
fof(f643, plain, ((e0 = op(e3, e0)) | ~ spl66_13), inference(avatar_component_clause, [], [f641])).
fof(f279, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4903, plain, (~ spl66_41 | spl66_46 | ~ spl66_140), inference(avatar_contradiction_clause, [], [f4902])).
fof(f4902, plain, ($false | (~ spl66_41 | spl66_46 | ~ spl66_140)), inference(subsumption_resolution, [], [f4898, f782])).
fof(f782, plain, (~ (e1 = op(e1, e0)) | spl66_46), inference(avatar_component_clause, [], [f781])).
fof(f781, plain, (spl66_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_46])])).
fof(f4898, plain, ((e1 = op(e1, e0)) | (~ spl66_41 | ~ spl66_140)), inference(backward_demodulation, [], [f1466, f762])).
fof(f1466, plain, ((e1 = op(e1, op(e1, e1))) | ~ spl66_140), inference(avatar_component_clause, [], [f1464])).
fof(f1464, plain, (spl66_140 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_140])])).
fof(f4891, plain, (~ spl66_36 | ~ spl66_48), inference(avatar_split_clause, [], [f4888, f789, f738])).
fof(f738, plain, (spl66_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_36])])).
fof(f4888, plain, (~ (e3 = op(e1, e3)) | ~ spl66_48), inference(backward_demodulation, [], [f267, f791])).
fof(f4890, plain, (~ spl66_16 | ~ spl66_48), inference(avatar_split_clause, [], [f4885, f789, f653])).
fof(f653, plain, (spl66_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_16])])).
fof(f4885, plain, (~ (e3 = op(e3, e0)) | ~ spl66_48), inference(backward_demodulation, [], [f239, f791])).
fof(f4883, plain, (~ spl66_1 | ~ spl66_49), inference(avatar_split_clause, [], [f4878, f794, f590])).
fof(f4878, plain, (~ (e0 = op(e3, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f255, f796])).
fof(f4875, plain, (~ spl66_51 | ~ spl66_55), inference(avatar_split_clause, [], [f4871, f819, f802])).
fof(f4871, plain, (~ (e2 = op(e0, e3)) | ~ spl66_55), inference(backward_demodulation, [], [f264, f821])).
fof(f821, plain, ((e2 = op(e0, e2)) | ~ spl66_55), inference(avatar_component_clause, [], [f819])).
fof(f264, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4874, plain, (~ spl66_39 | ~ spl66_55), inference(avatar_split_clause, [], [f4868, f819, f751])).
fof(f4868, plain, (~ (e2 = op(e1, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f247, f821])).
fof(f4865, plain, (spl66_21 | ~ spl66_31 | ~ spl66_135), inference(avatar_split_clause, [], [f4864, f1443, f717, f675])).
fof(f4864, plain, ((e0 = op(e2, e2)) | (~ spl66_31 | ~ spl66_135)), inference(forward_demodulation, [], [f1445, f719])).
fof(f4857, plain, (~ spl66_60 | ~ spl66_62 | spl66_162), inference(avatar_contradiction_clause, [], [f4856])).
fof(f4856, plain, ($false | (~ spl66_60 | ~ spl66_62 | spl66_162)), inference(subsumption_resolution, [], [f4855, f842])).
fof(f4855, plain, (~ (e3 = op(e0, e1)) | (~ spl66_62 | spl66_162)), inference(forward_demodulation, [], [f1571, f851])).
fof(f1571, plain, (~ (e3 = op(e0, op(e0, e0))) | spl66_162), inference(avatar_component_clause, [], [f1569])).
fof(f4849, plain, (~ spl66_43 | ~ spl66_11), inference(avatar_split_clause, [], [f4848, f632, f768])).
fof(f768, plain, (spl66_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_43])])).
fof(f4848, plain, (~ (e2 = op(e1, e1)) | ~ spl66_11), inference(forward_demodulation, [], [f245, f634])).
fof(f245, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4844, plain, (~ spl66_27 | ~ spl66_31), inference(avatar_split_clause, [], [f4843, f717, f700])).
fof(f700, plain, (spl66_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_27])])).
fof(f4843, plain, (~ (e2 = op(e2, e1)) | ~ spl66_31), inference(forward_demodulation, [], [f271, f719])).
fof(f4838, plain, (~ spl66_1 | ~ spl66_16 | spl66_134), inference(avatar_contradiction_clause, [], [f4837])).
fof(f4837, plain, ($false | (~ spl66_1 | ~ spl66_16 | spl66_134)), inference(subsumption_resolution, [], [f4831, f655])).
fof(f655, plain, ((e3 = op(e3, e0)) | ~ spl66_16), inference(avatar_component_clause, [], [f653])).
fof(f4831, plain, (~ (e3 = op(e3, e0)) | (~ spl66_1 | spl66_134)), inference(backward_demodulation, [], [f1439, f592])).
fof(f592, plain, ((e0 = op(e3, e3)) | ~ spl66_1), inference(avatar_component_clause, [], [f590])).
fof(f1439, plain, (~ (e3 = op(e3, op(e3, e3))) | spl66_134), inference(avatar_component_clause, [], [f1438])).
fof(f1438, plain, (spl66_134 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_134])])).
fof(f4822, plain, (spl66_1 | ~ spl66_16 | ~ spl66_131), inference(avatar_split_clause, [], [f4820, f1426, f653, f590])).
fof(f1426, plain, (spl66_131 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_131])])).
fof(f4820, plain, ((e0 = op(e3, e3)) | (~ spl66_16 | ~ spl66_131)), inference(backward_demodulation, [], [f1428, f655])).
fof(f1428, plain, ((e0 = op(e3, op(e3, e0))) | ~ spl66_131), inference(avatar_component_clause, [], [f1426])).
fof(f4821, plain, (~ spl66_8 | ~ spl66_16), inference(avatar_split_clause, [], [f4818, f653, f619])).
fof(f4818, plain, (~ (e3 = op(e3, e2)) | ~ spl66_16), inference(backward_demodulation, [], [f278, f655])).
fof(f4816, plain, (~ spl66_2 | ~ spl66_18), inference(avatar_split_clause, [], [f4812, f662, f594])).
fof(f4812, plain, (~ (e1 = op(e3, e3)) | ~ spl66_18), inference(backward_demodulation, [], [f258, f664])).
fof(f4809, plain, (~ spl66_20 | ~ spl66_24), inference(avatar_split_clause, [], [f4803, f687, f670])).
fof(f4803, plain, (~ (e3 = op(e2, e3)) | ~ spl66_24), inference(backward_demodulation, [], [f276, f689])).
fof(f4808, plain, (~ spl66_8 | ~ spl66_24), inference(avatar_split_clause, [], [f4802, f687, f619])).
fof(f4802, plain, (~ (e3 = op(e3, e2)) | ~ spl66_24), inference(backward_demodulation, [], [f252, f689])).
fof(f252, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4793, plain, (~ spl66_20 | ~ spl66_36), inference(avatar_split_clause, [], [f4790, f738, f670])).
fof(f4790, plain, (~ (e3 = op(e2, e3)) | ~ spl66_36), inference(backward_demodulation, [], [f256, f740])).
fof(f740, plain, ((e3 = op(e1, e3)) | ~ spl66_36), inference(avatar_component_clause, [], [f738])).
fof(f4787, plain, (~ spl66_35 | ~ spl66_39), inference(avatar_split_clause, [], [f4783, f751, f734])).
fof(f4783, plain, (~ (e2 = op(e1, e3)) | ~ spl66_39), inference(backward_demodulation, [], [f270, f753])).
fof(f270, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4763, plain, (~ spl66_53 | ~ spl66_62 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f4762])).
fof(f4762, plain, ($false | (~ spl66_53 | ~ spl66_62 | ~ spl66_145)), inference(subsumption_resolution, [], [f4761, f286])).
fof(f286, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f4761, plain, ((e1 = e2) | (~ spl66_53 | ~ spl66_62 | ~ spl66_145)), inference(forward_demodulation, [], [f4758, f851])).
fof(f4758, plain, ((op(e0, e0) = e2) | (~ spl66_53 | ~ spl66_145)), inference(backward_demodulation, [], [f1487, f813])).
fof(f4760, plain, (~ spl66_49 | ~ spl66_53), inference(avatar_split_clause, [], [f4754, f811, f794])).
fof(f4754, plain, (~ (e0 = op(e0, e3)) | ~ spl66_53), inference(backward_demodulation, [], [f264, f813])).
fof(f4748, plain, (~ spl66_35 | ~ spl66_60 | ~ spl66_62 | spl66_166), inference(avatar_split_clause, [], [f4747, f1589, f849, f840, f734])).
fof(f4747, plain, (~ (e2 = op(e1, e3)) | (~ spl66_60 | ~ spl66_62 | spl66_166)), inference(forward_demodulation, [], [f4742, f842])).
fof(f4742, plain, (~ (e2 = op(e1, op(e0, e1))) | (~ spl66_62 | spl66_166)), inference(backward_demodulation, [], [f1591, f851])).
fof(f4745, plain, (~ spl66_54 | ~ spl66_62), inference(avatar_split_clause, [], [f4735, f849, f815])).
fof(f4735, plain, (~ (e1 = op(e0, e2)) | ~ spl66_62), inference(backward_demodulation, [], [f260, f851])).
fof(f260, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f4731, plain, (spl66_6 | ~ spl66_11 | ~ spl66_132), inference(avatar_split_clause, [], [f4730, f1430, f632, f611])).
fof(f4730, plain, ((e1 = op(e3, e2)) | (~ spl66_11 | ~ spl66_132)), inference(forward_demodulation, [], [f1432, f634])).
fof(f1432, plain, ((e1 = op(e3, op(e3, e1))) | ~ spl66_132), inference(avatar_component_clause, [], [f1430])).
fof(f4727, plain, (~ spl66_42 | spl66_140), inference(avatar_contradiction_clause, [], [f4726])).
fof(f4726, plain, ($false | (~ spl66_42 | spl66_140)), inference(subsumption_resolution, [], [f4725, f766])).
fof(f4725, plain, (~ (e1 = op(e1, e1)) | (~ spl66_42 | spl66_140)), inference(forward_demodulation, [], [f1465, f766])).
fof(f1465, plain, (~ (e1 = op(e1, op(e1, e1))) | spl66_140), inference(avatar_component_clause, [], [f1464])).
fof(f4722, plain, (~ spl66_137 | ~ spl66_150), inference(avatar_contradiction_clause, [], [f4721])).
fof(f4721, plain, ($false | (~ spl66_137 | ~ spl66_150)), inference(subsumption_resolution, [], [f4720, f286])).
fof(f4720, plain, ((e1 = e2) | (~ spl66_137 | ~ spl66_150)), inference(backward_demodulation, [], [f1453, f1512])).
fof(f1512, plain, ((e1 = op(e2, op(e2, e2))) | ~ spl66_150), inference(avatar_component_clause, [], [f1511])).
fof(f1511, plain, (spl66_150 <=> (e1 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_150])])).
fof(f4716, plain, (~ spl66_3 | ~ spl66_11), inference(avatar_split_clause, [], [f4715, f632, f598])).
fof(f598, plain, (spl66_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_3])])).
fof(f4715, plain, (~ (e2 = op(e3, e3)) | ~ spl66_11), inference(forward_demodulation, [], [f281, f634])).
fof(f281, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4710, plain, (~ spl66_7 | ~ spl66_11), inference(avatar_split_clause, [], [f4705, f632, f615])).
fof(f615, plain, (spl66_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_7])])).
fof(f4705, plain, (~ (e2 = op(e3, e2)) | ~ spl66_11), inference(backward_demodulation, [], [f280, f634])).
fof(f280, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4704, plain, (~ spl66_9 | ~ spl66_13), inference(avatar_split_clause, [], [f4698, f641, f624])).
fof(f4698, plain, (~ (e0 = op(e3, e1)) | ~ spl66_13), inference(backward_demodulation, [], [f277, f643])).
fof(f4696, plain, (~ spl66_7 | ~ spl66_23), inference(avatar_split_clause, [], [f4691, f683, f615])).
fof(f4691, plain, (~ (e2 = op(e3, e2)) | ~ spl66_23), inference(backward_demodulation, [], [f252, f685])).
fof(f4683, plain, (spl66_25 | ~ spl66_30 | ~ spl66_135), inference(avatar_split_clause, [], [f4681, f1443, f713, f692])).
fof(f4681, plain, ((e0 = op(e2, e1)) | (~ spl66_30 | ~ spl66_135)), inference(backward_demodulation, [], [f1445, f715])).
fof(f4672, plain, (~ spl66_33 | ~ spl66_37), inference(avatar_split_clause, [], [f4667, f743, f726])).
fof(f726, plain, (spl66_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_33])])).
fof(f4667, plain, (~ (e0 = op(e1, e3)) | ~ spl66_37), inference(backward_demodulation, [], [f270, f745])).
fof(f4655, plain, (~ spl66_33 | ~ spl66_49), inference(avatar_split_clause, [], [f4650, f794, f726])).
fof(f4650, plain, (~ (e0 = op(e1, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f253, f796])).
fof(f4649, plain, (~ spl66_22 | ~ spl66_54), inference(avatar_split_clause, [], [f4643, f815, f679])).
fof(f4643, plain, (~ (e1 = op(e2, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f248, f817])).
fof(f248, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4638, plain, (~ spl66_51 | ~ spl66_63), inference(avatar_split_clause, [], [f4630, f853, f802])).
fof(f4630, plain, (~ (e2 = op(e0, e3)) | ~ spl66_63), inference(backward_demodulation, [], [f261, f855])).
fof(f4637, plain, (~ spl66_47 | ~ spl66_63), inference(avatar_split_clause, [], [f4626, f853, f785])).
fof(f4626, plain, (~ (e2 = op(e1, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f235, f855])).
fof(f4625, plain, (spl66_50 | ~ spl66_60 | ~ spl66_144), inference(avatar_contradiction_clause, [], [f4624])).
fof(f4624, plain, ($false | (spl66_50 | ~ spl66_60 | ~ spl66_144)), inference(subsumption_resolution, [], [f4623, f799])).
fof(f799, plain, (~ (e1 = op(e0, e3)) | spl66_50), inference(avatar_component_clause, [], [f798])).
fof(f4623, plain, ((e1 = op(e0, e3)) | (~ spl66_60 | ~ spl66_144)), inference(forward_demodulation, [], [f1483, f842])).
fof(f1483, plain, ((e1 = op(e0, op(e0, e1))) | ~ spl66_144), inference(avatar_component_clause, [], [f1481])).
fof(f4622, plain, (~ spl66_11 | ~ spl66_2 | spl66_151), inference(avatar_split_clause, [], [f4621, f1516, f594, f632])).
fof(f4621, plain, (~ (e2 = op(e3, e1)) | (~ spl66_2 | spl66_151)), inference(forward_demodulation, [], [f1518, f596])).
fof(f1518, plain, (~ (e2 = op(e3, op(e3, e3))) | spl66_151), inference(avatar_component_clause, [], [f1516])).
fof(f4620, plain, (~ spl66_9 | ~ spl66_2 | spl66_157), inference(avatar_split_clause, [], [f4619, f1544, f594, f624])).
fof(f1544, plain, (spl66_157 <=> (e0 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_157])])).
fof(f4619, plain, (~ (e0 = op(e3, e1)) | (~ spl66_2 | spl66_157)), inference(forward_demodulation, [], [f1546, f596])).
fof(f1546, plain, (~ (e0 = op(e3, op(e3, e3))) | spl66_157), inference(avatar_component_clause, [], [f1544])).
fof(f4602, plain, (~ spl66_2 | ~ spl66_9 | ~ spl66_47 | spl66_163), inference(avatar_contradiction_clause, [], [f4601])).
fof(f4601, plain, ($false | (~ spl66_2 | ~ spl66_9 | ~ spl66_47 | spl66_163)), inference(subsumption_resolution, [], [f4600, f787])).
fof(f4600, plain, (~ (e2 = op(e1, e0)) | (~ spl66_2 | ~ spl66_9 | spl66_163)), inference(backward_demodulation, [], [f4512, f626])).
fof(f626, plain, ((e0 = op(e3, e1)) | ~ spl66_9), inference(avatar_component_clause, [], [f624])).
fof(f4512, plain, (~ (e2 = op(e1, op(e3, e1))) | (~ spl66_2 | spl66_163)), inference(backward_demodulation, [], [f1576, f596])).
fof(f1576, plain, (~ (e2 = op(op(e3, e3), op(e3, op(e3, e3)))) | spl66_163), inference(avatar_component_clause, [], [f1574])).
fof(f1574, plain, (spl66_163 <=> (e2 = op(op(e3, e3), op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl66_163])])).
fof(f4593, plain, (~ spl66_12 | ~ spl66_16), inference(avatar_split_clause, [], [f4592, f653, f636])).
fof(f4592, plain, (~ (e3 = op(e3, e1)) | ~ spl66_16), inference(backward_demodulation, [], [f277, f655])).
fof(f4589, plain, (~ spl66_19 | ~ spl66_27), inference(avatar_split_clause, [], [f4585, f700, f666])).
fof(f4585, plain, (~ (e2 = op(e2, e3)) | ~ spl66_27), inference(backward_demodulation, [], [f275, f702])).
fof(f702, plain, ((e2 = op(e2, e1)) | ~ spl66_27), inference(avatar_component_clause, [], [f700])).
fof(f4583, plain, (~ spl66_25 | ~ spl66_29), inference(avatar_split_clause, [], [f4576, f709, f692])).
fof(f4576, plain, (~ (e0 = op(e2, e1)) | ~ spl66_29), inference(backward_demodulation, [], [f271, f711])).
fof(f4582, plain, (~ spl66_13 | ~ spl66_29), inference(avatar_split_clause, [], [f4575, f709, f641])).
fof(f4575, plain, (~ (e0 = op(e3, e0)) | ~ spl66_29), inference(backward_demodulation, [], [f240, f711])).
fof(f240, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f4574, plain, (~ spl66_51 | ~ spl66_53 | ~ spl66_146), inference(avatar_contradiction_clause, [], [f4573])).
fof(f4573, plain, ($false | (~ spl66_51 | ~ spl66_53 | ~ spl66_146)), inference(subsumption_resolution, [], [f4572, f285])).
fof(f4572, plain, ((e0 = e3) | (~ spl66_51 | ~ spl66_53 | ~ spl66_146)), inference(forward_demodulation, [], [f4570, f813])).
fof(f4570, plain, ((e3 = op(e0, e2)) | (~ spl66_51 | ~ spl66_146)), inference(backward_demodulation, [], [f1491, f804])).
fof(f4571, plain, (~ spl66_19 | ~ spl66_51), inference(avatar_split_clause, [], [f4568, f802, f666])).
fof(f4568, plain, (~ (e2 = op(e2, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f254, f804])).
fof(f4565, plain, (~ spl66_52 | ~ spl66_60), inference(avatar_split_clause, [], [f4560, f840, f806])).
fof(f4560, plain, (~ (e3 = op(e0, e3)) | ~ spl66_60), inference(backward_demodulation, [], [f263, f842])).
fof(f4564, plain, (~ spl66_12 | ~ spl66_60), inference(avatar_split_clause, [], [f4559, f840, f636])).
fof(f4559, plain, (~ (e3 = op(e3, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f243, f842])).
fof(f4552, plain, (~ spl66_25 | ~ spl66_22 | spl66_159), inference(avatar_split_clause, [], [f4551, f1553, f679, f692])).
fof(f4551, plain, (~ (e0 = op(e2, e1)) | (~ spl66_22 | spl66_159)), inference(forward_demodulation, [], [f1555, f681])).
fof(f681, plain, ((e1 = op(e2, e2)) | ~ spl66_22), inference(avatar_component_clause, [], [f679])).
fof(f4550, plain, (~ spl66_59 | ~ spl66_62 | spl66_161), inference(avatar_split_clause, [], [f4549, f1563, f849, f836])).
fof(f4549, plain, (~ (e2 = op(e0, e1)) | (~ spl66_62 | spl66_161)), inference(forward_demodulation, [], [f1565, f851])).
fof(f1565, plain, (~ (e2 = op(e0, op(e0, e0))) | spl66_161), inference(avatar_component_clause, [], [f1563])).
fof(f4545, plain, (spl66_27 | ~ spl66_22 | ~ spl66_137), inference(avatar_split_clause, [], [f4544, f1451, f679, f700])).
fof(f4544, plain, ((e2 = op(e2, e1)) | (~ spl66_22 | ~ spl66_137)), inference(forward_demodulation, [], [f1453, f681])).
fof(f4505, plain, (~ spl66_3 | ~ spl66_7), inference(avatar_split_clause, [], [f4502, f615, f598])).
fof(f4502, plain, (~ (e2 = op(e3, e3)) | ~ spl66_7), inference(backward_demodulation, [], [f282, f617])).
fof(f617, plain, ((e2 = op(e3, e2)) | ~ spl66_7), inference(avatar_component_clause, [], [f615])).
fof(f4501, plain, (~ spl66_5 | ~ spl66_13), inference(avatar_split_clause, [], [f4495, f641, f607])).
fof(f4495, plain, (~ (e0 = op(e3, e2)) | ~ spl66_13), inference(backward_demodulation, [], [f278, f643])).
fof(f4493, plain, (~ spl66_3 | ~ spl66_19), inference(avatar_split_clause, [], [f4490, f666, f598])).
fof(f4490, plain, (~ (e2 = op(e3, e3)) | ~ spl66_19), inference(backward_demodulation, [], [f258, f668])).
fof(f668, plain, ((e2 = op(e2, e3)) | ~ spl66_19), inference(avatar_component_clause, [], [f666])).
fof(f4480, plain, (~ spl66_17 | ~ spl66_25), inference(avatar_split_clause, [], [f4475, f692, f658])).
fof(f4475, plain, (~ (e0 = op(e2, e3)) | ~ spl66_25), inference(backward_demodulation, [], [f275, f694])).
fof(f4472, plain, (~ spl66_17 | ~ spl66_33), inference(avatar_split_clause, [], [f4466, f726, f658])).
fof(f4466, plain, (~ (e0 = op(e2, e3)) | ~ spl66_33), inference(backward_demodulation, [], [f256, f728])).
fof(f728, plain, ((e0 = op(e1, e3)) | ~ spl66_33), inference(avatar_component_clause, [], [f726])).
fof(f4464, plain, (~ spl66_34 | ~ spl66_42), inference(avatar_split_clause, [], [f4456, f764, f730])).
fof(f4456, plain, (~ (e1 = op(e1, e3)) | ~ spl66_42), inference(backward_demodulation, [], [f269, f766])).
fof(f269, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4454, plain, (~ spl66_43 | ~ spl66_47), inference(avatar_split_clause, [], [f4450, f785, f768])).
fof(f4450, plain, (~ (e2 = op(e1, e1)) | ~ spl66_47), inference(backward_demodulation, [], [f265, f787])).
fof(f265, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f4447, plain, (~ spl66_5 | ~ spl66_53), inference(avatar_split_clause, [], [f4442, f811, f607])).
fof(f4442, plain, (~ (e0 = op(e3, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f249, f813])).
fof(f4439, plain, (~ spl66_43 | ~ spl66_59), inference(avatar_split_clause, [], [f4434, f836, f768])).
fof(f4434, plain, (~ (e2 = op(e1, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f241, f838])).
fof(f4428, plain, (~ spl66_14 | ~ spl66_62), inference(avatar_split_clause, [], [f4419, f849, f645])).
fof(f4419, plain, (~ (e1 = op(e3, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f237, f851])).
fof(f4417, plain, (spl66_2 | ~ spl66_12 | ~ spl66_132), inference(avatar_split_clause, [], [f4416, f1430, f636, f594])).
fof(f4416, plain, ((e1 = op(e3, e3)) | (~ spl66_12 | ~ spl66_132)), inference(forward_demodulation, [], [f1432, f638])).
fof(f638, plain, ((e3 = op(e3, e1)) | ~ spl66_12), inference(avatar_component_clause, [], [f636])).
fof(f4415, plain, (~ spl66_17 | ~ spl66_32 | spl66_135), inference(avatar_split_clause, [], [f4414, f1443, f721, f658])).
fof(f4414, plain, (~ (e0 = op(e2, e3)) | (~ spl66_32 | spl66_135)), inference(forward_demodulation, [], [f1444, f723])).
fof(f1444, plain, (~ (e0 = op(e2, op(e2, e0))) | spl66_135), inference(avatar_component_clause, [], [f1443])).
fof(f4390, plain, (~ spl66_24 | ~ spl66_40), inference(avatar_split_clause, [], [f4389, f755, f687])).
fof(f4389, plain, (~ (e3 = op(e2, e2)) | ~ spl66_40), inference(forward_demodulation, [], [f250, f757])).
fof(f4386, plain, (~ spl66_24 | ~ spl66_32), inference(avatar_split_clause, [], [f4385, f721, f687])).
fof(f4385, plain, (~ (e3 = op(e2, e2)) | ~ spl66_32), inference(forward_demodulation, [], [f272, f723])).
fof(f4383, plain, (~ spl66_4 | ~ spl66_12), inference(avatar_split_clause, [], [f4382, f636, f602])).
fof(f4382, plain, (~ (e3 = op(e3, e3)) | ~ spl66_12), inference(forward_demodulation, [], [f281, f638])).
fof(f4381, plain, (~ spl66_4 | ~ spl66_52), inference(avatar_split_clause, [], [f4380, f806, f602])).
fof(f4380, plain, (~ (e3 = op(e3, e3)) | ~ spl66_52), inference(forward_demodulation, [], [f255, f808])).
fof(f4352, plain, (~ spl66_16 | ~ spl66_32), inference(avatar_split_clause, [], [f4348, f721, f653])).
fof(f4348, plain, (~ (e3 = op(e3, e0)) | ~ spl66_32), inference(backward_demodulation, [], [f240, f723])).
fof(f4347, plain, (~ spl66_34 | ~ spl66_43 | ~ spl66_142), inference(avatar_contradiction_clause, [], [f4346])).
fof(f4346, plain, ($false | (~ spl66_34 | ~ spl66_43 | ~ spl66_142)), inference(subsumption_resolution, [], [f4345, f288])).
fof(f4345, plain, ((e2 = e3) | (~ spl66_34 | ~ spl66_43 | ~ spl66_142)), inference(forward_demodulation, [], [f4343, f770])).
fof(f770, plain, ((e2 = op(e1, e1)) | ~ spl66_43), inference(avatar_component_clause, [], [f768])).
fof(f4343, plain, ((e3 = op(e1, e1)) | (~ spl66_34 | ~ spl66_142)), inference(backward_demodulation, [], [f1474, f732])).
fof(f4339, plain, (~ spl66_17 | ~ spl66_40 | ~ spl66_43 | spl66_152), inference(avatar_split_clause, [], [f4335, f1521, f768, f755, f658])).
fof(f4335, plain, (~ (e0 = op(e2, e3)) | (~ spl66_40 | ~ spl66_43 | spl66_152)), inference(backward_demodulation, [], [f4210, f757])).
fof(f4210, plain, (~ (e0 = op(e2, op(e1, e2))) | (~ spl66_43 | spl66_152)), inference(backward_demodulation, [], [f1523, f770])).
fof(f4338, plain, (~ spl66_36 | ~ spl66_40), inference(avatar_split_clause, [], [f4333, f755, f738])).
fof(f4333, plain, (~ (e3 = op(e1, e3)) | ~ spl66_40), inference(backward_demodulation, [], [f270, f757])).
fof(f4331, plain, (~ spl66_45 | spl66_139), inference(avatar_contradiction_clause, [], [f4330])).
fof(f4330, plain, ($false | (~ spl66_45 | spl66_139)), inference(subsumption_resolution, [], [f4327, f779])).
fof(f4327, plain, (~ (e0 = op(e1, e0)) | (~ spl66_45 | spl66_139)), inference(backward_demodulation, [], [f1461, f779])).
fof(f1461, plain, (~ (e0 = op(e1, op(e1, e0))) | spl66_139), inference(avatar_component_clause, [], [f1460])).
fof(f1460, plain, (spl66_139 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_139])])).
fof(f4319, plain, (~ spl66_36 | ~ spl66_52), inference(avatar_split_clause, [], [f4316, f806, f738])).
fof(f4316, plain, (~ (e3 = op(e1, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f253, f808])).
fof(f4313, plain, (~ spl66_6 | ~ spl66_54), inference(avatar_split_clause, [], [f4307, f815, f611])).
fof(f4307, plain, (~ (e1 = op(e3, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f249, f817])).
fof(f4304, plain, (~ spl66_9 | ~ spl66_57), inference(avatar_split_clause, [], [f4297, f828, f624])).
fof(f4297, plain, (~ (e0 = op(e3, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f243, f830])).
fof(f4295, plain, (~ spl66_6 | ~ spl66_3 | spl66_148), inference(avatar_split_clause, [], [f4294, f1502, f598, f611])).
fof(f1502, plain, (spl66_148 <=> (e1 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_148])])).
fof(f4294, plain, (~ (e1 = op(e3, e2)) | (~ spl66_3 | spl66_148)), inference(forward_demodulation, [], [f1504, f600])).
fof(f600, plain, ((e2 = op(e3, e3)) | ~ spl66_3), inference(avatar_component_clause, [], [f598])).
fof(f1504, plain, (~ (e1 = op(e3, op(e3, e3))) | spl66_148), inference(avatar_component_clause, [], [f1502])).
fof(f4282, plain, (~ spl66_3 | spl66_8 | ~ spl66_134), inference(avatar_contradiction_clause, [], [f4281])).
fof(f4281, plain, ($false | (~ spl66_3 | spl66_8 | ~ spl66_134)), inference(subsumption_resolution, [], [f4280, f620])).
fof(f620, plain, (~ (e3 = op(e3, e2)) | spl66_8), inference(avatar_component_clause, [], [f619])).
fof(f4280, plain, ((e3 = op(e3, e2)) | (~ spl66_3 | ~ spl66_134)), inference(forward_demodulation, [], [f1440, f600])).
fof(f1440, plain, ((e3 = op(e3, op(e3, e3))) | ~ spl66_134), inference(avatar_component_clause, [], [f1438])).
fof(f4278, plain, (~ spl66_51 | ~ spl66_3), inference(avatar_split_clause, [], [f4277, f598, f802])).
fof(f4277, plain, (~ (e2 = op(e0, e3)) | ~ spl66_3), inference(forward_demodulation, [], [f255, f600])).
fof(f4252, plain, (~ spl66_5 | ~ spl66_9), inference(avatar_split_clause, [], [f4246, f624, f607])).
fof(f4246, plain, (~ (e0 = op(e3, e2)) | ~ spl66_9), inference(backward_demodulation, [], [f280, f626])).
fof(f4240, plain, (spl66_18 | ~ spl66_28 | ~ spl66_136), inference(avatar_split_clause, [], [f4238, f1447, f704, f662])).
fof(f1447, plain, (spl66_136 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_136])])).
fof(f4238, plain, ((e1 = op(e2, e3)) | (~ spl66_28 | ~ spl66_136)), inference(backward_demodulation, [], [f1449, f706])).
fof(f1449, plain, ((e1 = op(e2, op(e2, e1))) | ~ spl66_136), inference(avatar_component_clause, [], [f1447])).
fof(f4227, plain, (~ spl66_5 | ~ spl66_37), inference(avatar_split_clause, [], [f4219, f743, f607])).
fof(f4219, plain, (~ (e0 = op(e3, e2)) | ~ spl66_37), inference(backward_demodulation, [], [f251, f745])).
fof(f4218, plain, (~ spl66_37 | ~ spl66_43 | spl66_165), inference(avatar_split_clause, [], [f4214, f1583, f768, f743])).
fof(f1583, plain, (spl66_165 <=> (e0 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_165])])).
fof(f4214, plain, (~ (e0 = op(e1, e2)) | (~ spl66_43 | spl66_165)), inference(backward_demodulation, [], [f1585, f770])).
fof(f1585, plain, (~ (e0 = op(e1, op(e1, e1))) | spl66_165), inference(avatar_component_clause, [], [f1583])).
fof(f4217, plain, (~ spl66_40 | ~ spl66_43 | spl66_155), inference(avatar_split_clause, [], [f4212, f1535, f768, f755])).
fof(f1535, plain, (spl66_155 <=> (e3 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_155])])).
fof(f4212, plain, (~ (e3 = op(e1, e2)) | (~ spl66_43 | spl66_155)), inference(backward_demodulation, [], [f1537, f770])).
fof(f1537, plain, (~ (e3 = op(e1, op(e1, e1))) | spl66_155), inference(avatar_component_clause, [], [f1535])).
fof(f4216, plain, (~ spl66_35 | ~ spl66_43), inference(avatar_split_clause, [], [f4207, f768, f734])).
fof(f4207, plain, (~ (e2 = op(e1, e3)) | ~ spl66_43), inference(backward_demodulation, [], [f269, f770])).
fof(f4197, plain, (~ spl66_54 | ~ spl66_58), inference(avatar_split_clause, [], [f4193, f832, f815])).
fof(f4193, plain, (~ (e1 = op(e0, e2)) | ~ spl66_58), inference(backward_demodulation, [], [f262, f834])).
fof(f4196, plain, (~ spl66_26 | ~ spl66_58), inference(avatar_split_clause, [], [f4191, f832, f696])).
fof(f4191, plain, (~ (e1 = op(e2, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f242, f834])).
fof(f242, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4176, plain, (~ spl66_56 | ~ spl66_63 | spl66_162), inference(avatar_split_clause, [], [f4175, f1569, f853, f823])).
fof(f4175, plain, (~ (e3 = op(e0, e2)) | (~ spl66_63 | spl66_162)), inference(forward_demodulation, [], [f1571, f855])).
fof(f4172, plain, (~ spl66_55 | ~ spl66_23), inference(avatar_split_clause, [], [f4171, f683, f819])).
fof(f4171, plain, (~ (e2 = op(e0, e2)) | ~ spl66_23), inference(forward_demodulation, [], [f248, f685])).
fof(f4170, plain, (~ spl66_55 | ~ spl66_63), inference(avatar_split_clause, [], [f4169, f853, f819])).
fof(f4169, plain, (~ (e2 = op(e0, e2)) | ~ spl66_63), inference(forward_demodulation, [], [f260, f855])).
fof(f4167, plain, (~ spl66_27 | ~ spl66_23), inference(avatar_split_clause, [], [f4166, f683, f700])).
fof(f4166, plain, (~ (e2 = op(e2, e1)) | ~ spl66_23), inference(forward_demodulation, [], [f274, f685])).
fof(f274, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4163, plain, (~ spl66_4 | ~ spl66_16), inference(avatar_split_clause, [], [f4162, f653, f602])).
fof(f4162, plain, (~ (e3 = op(e3, e3)) | ~ spl66_16), inference(forward_demodulation, [], [f279, f655])).
fof(f4146, plain, (~ spl66_20 | spl66_138), inference(avatar_contradiction_clause, [], [f4145])).
fof(f4145, plain, ($false | (~ spl66_20 | spl66_138)), inference(subsumption_resolution, [], [f4144, f672])).
fof(f672, plain, ((e3 = op(e2, e3)) | ~ spl66_20), inference(avatar_component_clause, [], [f670])).
fof(f4144, plain, (~ (e3 = op(e2, e3)) | (~ spl66_20 | spl66_138)), inference(backward_demodulation, [], [f1456, f672])).
fof(f4143, plain, (~ spl66_21 | ~ spl66_23 | spl66_159), inference(avatar_split_clause, [], [f4136, f1553, f683, f675])).
fof(f4136, plain, (~ (e0 = op(e2, e2)) | (~ spl66_23 | spl66_159)), inference(backward_demodulation, [], [f1555, f685])).
fof(f4139, plain, (~ spl66_23 | spl66_137), inference(avatar_contradiction_clause, [], [f4138])).
fof(f4138, plain, ($false | (~ spl66_23 | spl66_137)), inference(subsumption_resolution, [], [f4134, f685])).
fof(f4134, plain, (~ (e2 = op(e2, e2)) | (~ spl66_23 | spl66_137)), inference(backward_demodulation, [], [f1452, f685])).
fof(f1452, plain, (~ (e2 = op(e2, op(e2, e2))) | spl66_137), inference(avatar_component_clause, [], [f1451])).
fof(f4130, plain, (~ spl66_21 | ~ spl66_29), inference(avatar_split_clause, [], [f4124, f709, f675])).
fof(f4124, plain, (~ (e0 = op(e2, e2)) | ~ spl66_29), inference(backward_demodulation, [], [f272, f711])).
fof(f4120, plain, (~ spl66_19 | ~ spl66_35), inference(avatar_split_clause, [], [f4117, f734, f666])).
fof(f4117, plain, (~ (e2 = op(e2, e3)) | ~ spl66_35), inference(backward_demodulation, [], [f256, f736])).
fof(f4116, plain, (spl66_35 | ~ spl66_40 | ~ spl66_141), inference(avatar_split_clause, [], [f4114, f1468, f755, f734])).
fof(f4114, plain, ((e2 = op(e1, e3)) | (~ spl66_40 | ~ spl66_141)), inference(backward_demodulation, [], [f1470, f757])).
fof(f1470, plain, ((e2 = op(e1, op(e1, e2))) | ~ spl66_141), inference(avatar_component_clause, [], [f1468])).
fof(f4110, plain, (~ spl66_33 | ~ spl66_41), inference(avatar_split_clause, [], [f4101, f760, f726])).
fof(f4101, plain, (~ (e0 = op(e1, e3)) | ~ spl66_41), inference(backward_demodulation, [], [f269, f762])).
fof(f4109, plain, (~ spl66_9 | ~ spl66_41), inference(avatar_split_clause, [], [f4099, f760, f624])).
fof(f4099, plain, (~ (e0 = op(e3, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f245, f762])).
fof(f4075, plain, (spl66_41 | ~ spl66_46 | ~ spl66_139), inference(avatar_split_clause, [], [f4074, f1460, f781, f760])).
fof(f4074, plain, ((e0 = op(e1, e1)) | (~ spl66_46 | ~ spl66_139)), inference(forward_demodulation, [], [f1462, f783])).
fof(f783, plain, ((e1 = op(e1, e0)) | ~ spl66_46), inference(avatar_component_clause, [], [f781])).
fof(f1462, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl66_139), inference(avatar_component_clause, [], [f1460])).
fof(f4055, plain, (spl66_59 | ~ spl66_54 | ~ spl66_145), inference(avatar_split_clause, [], [f4003, f1485, f815, f836])).
fof(f4003, plain, ((e2 = op(e0, e1)) | (~ spl66_54 | ~ spl66_145)), inference(backward_demodulation, [], [f1487, f817])).
fof(f4051, plain, (~ spl66_34 | ~ spl66_46), inference(avatar_split_clause, [], [f4050, f781, f730])).
fof(f4050, plain, (~ (e1 = op(e1, e3)) | ~ spl66_46), inference(forward_demodulation, [], [f267, f783])).
fof(f4048, plain, (~ spl66_22 | ~ spl66_26), inference(avatar_split_clause, [], [f4047, f696, f679])).
fof(f4047, plain, (~ (e1 = op(e2, e2)) | ~ spl66_26), inference(forward_demodulation, [], [f274, f698])).
fof(f4043, plain, (~ spl66_10 | ~ spl66_26), inference(avatar_split_clause, [], [f4042, f696, f628])).
fof(f628, plain, (spl66_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_10])])).
fof(f4042, plain, (~ (e1 = op(e3, e1)) | ~ spl66_26), inference(forward_demodulation, [], [f246, f698])).
fof(f4041, plain, (spl66_12 | ~ spl66_2 | ~ spl66_134), inference(avatar_split_clause, [], [f3985, f1438, f594, f636])).
fof(f3985, plain, ((e3 = op(e3, e1)) | (~ spl66_2 | ~ spl66_134)), inference(forward_demodulation, [], [f1440, f596])).
fof(f4038, plain, (~ spl66_7 | ~ spl66_15), inference(avatar_split_clause, [], [f4035, f649, f615])).
fof(f4035, plain, (~ (e2 = op(e3, e2)) | ~ spl66_15), inference(backward_demodulation, [], [f278, f651])).
fof(f4019, plain, (~ spl66_7 | ~ spl66_39), inference(avatar_split_clause, [], [f4016, f751, f615])).
fof(f4016, plain, (~ (e2 = op(e3, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f251, f753])).
fof(f4015, plain, (~ spl66_38 | ~ spl66_46), inference(avatar_split_clause, [], [f4010, f781, f747])).
fof(f4010, plain, (~ (e1 = op(e1, e2)) | ~ spl66_46), inference(backward_demodulation, [], [f266, f783])).
fof(f4014, plain, (~ spl66_30 | ~ spl66_46), inference(avatar_split_clause, [], [f4008, f781, f713])).
fof(f4008, plain, (~ (e1 = op(e2, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f238, f783])).
fof(f3997, plain, (spl66_54 | ~ spl66_59 | ~ spl66_144), inference(avatar_split_clause, [], [f3994, f1481, f836, f815])).
fof(f3994, plain, ((e1 = op(e0, e2)) | (~ spl66_59 | ~ spl66_144)), inference(backward_demodulation, [], [f1483, f838])).
fof(f3996, plain, (~ spl66_51 | ~ spl66_59), inference(avatar_split_clause, [], [f3992, f836, f802])).
fof(f3992, plain, (~ (e2 = op(e0, e3)) | ~ spl66_59), inference(backward_demodulation, [], [f263, f838])).
fof(f3995, plain, (~ spl66_27 | ~ spl66_59), inference(avatar_split_clause, [], [f3990, f836, f700])).
fof(f3990, plain, (~ (e2 = op(e2, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f242, f838])).
fof(f3983, plain, (~ spl66_30 | ~ spl66_21 | spl66_150), inference(avatar_split_clause, [], [f3982, f1511, f675, f713])).
fof(f3982, plain, (~ (e1 = op(e2, e0)) | (~ spl66_21 | spl66_150)), inference(forward_demodulation, [], [f1513, f677])).
fof(f677, plain, ((e0 = op(e2, e2)) | ~ spl66_21), inference(avatar_component_clause, [], [f675])).
fof(f1513, plain, (~ (e1 = op(e2, op(e2, e2))) | spl66_150), inference(avatar_component_clause, [], [f1511])).
fof(f3981, plain, (spl66_32 | ~ spl66_21 | ~ spl66_154), inference(avatar_split_clause, [], [f3980, f1530, f675, f721])).
fof(f1530, plain, (spl66_154 <=> (e3 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_154])])).
fof(f3980, plain, ((e3 = op(e2, e0)) | (~ spl66_21 | ~ spl66_154)), inference(forward_demodulation, [], [f1531, f677])).
fof(f1531, plain, ((e3 = op(e2, op(e2, e2))) | ~ spl66_154), inference(avatar_component_clause, [], [f1530])).
fof(f3966, plain, (~ spl66_40 | ~ spl66_44), inference(avatar_split_clause, [], [f3965, f772, f755])).
fof(f3965, plain, (~ (e3 = op(e1, e2)) | ~ spl66_44), inference(forward_demodulation, [], [f268, f774])).
fof(f3958, plain, (~ spl66_15 | ~ spl66_44 | spl66_164 | ~ spl66_165), inference(avatar_split_clause, [], [f3909, f1583, f1579, f772, f649])).
fof(f1579, plain, (spl66_164 <=> (e2 = op(op(e1, e1), op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl66_164])])).
fof(f3909, plain, (~ (e2 = op(e3, e0)) | (~ spl66_44 | spl66_164 | ~ spl66_165)), inference(backward_demodulation, [], [f3879, f774])).
fof(f3879, plain, (~ (e2 = op(op(e1, e1), e0)) | (spl66_164 | ~ spl66_165)), inference(backward_demodulation, [], [f1581, f1584])).
fof(f1584, plain, ((e0 = op(e1, op(e1, e1))) | ~ spl66_165), inference(avatar_component_clause, [], [f1583])).
fof(f1581, plain, (~ (e2 = op(op(e1, e1), op(e1, op(e1, e1)))) | spl66_164), inference(avatar_component_clause, [], [f1579])).
fof(f3940, plain, (~ spl66_9 | spl66_14 | ~ spl66_132), inference(avatar_contradiction_clause, [], [f3939])).
fof(f3939, plain, ($false | (~ spl66_9 | spl66_14 | ~ spl66_132)), inference(subsumption_resolution, [], [f3937, f646])).
fof(f646, plain, (~ (e1 = op(e3, e0)) | spl66_14), inference(avatar_component_clause, [], [f645])).
fof(f3937, plain, ((e1 = op(e3, e0)) | (~ spl66_9 | ~ spl66_132)), inference(backward_demodulation, [], [f1432, f626])).
fof(f3938, plain, (~ spl66_1 | ~ spl66_9), inference(avatar_split_clause, [], [f3933, f624, f590])).
fof(f3933, plain, (~ (e0 = op(e3, e3)) | ~ spl66_9), inference(backward_demodulation, [], [f281, f626])).
fof(f3924, plain, (~ spl66_33 | ~ spl66_47 | ~ spl66_142), inference(avatar_contradiction_clause, [], [f3923])).
fof(f3923, plain, ($false | (~ spl66_33 | ~ spl66_47 | ~ spl66_142)), inference(subsumption_resolution, [], [f3922, f288])).
fof(f3922, plain, ((e2 = e3) | (~ spl66_33 | ~ spl66_47 | ~ spl66_142)), inference(forward_demodulation, [], [f3920, f787])).
fof(f3920, plain, ((e3 = op(e1, e0)) | (~ spl66_33 | ~ spl66_142)), inference(backward_demodulation, [], [f1474, f728])).
fof(f3921, plain, (~ spl66_1 | ~ spl66_33), inference(avatar_split_clause, [], [f3916, f726, f590])).
fof(f3916, plain, (~ (e0 = op(e3, e3)) | ~ spl66_33), inference(backward_demodulation, [], [f257, f728])).
fof(f257, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3911, plain, (~ spl66_36 | ~ spl66_44), inference(avatar_split_clause, [], [f3905, f772, f738])).
fof(f3905, plain, (~ (e3 = op(e1, e3)) | ~ spl66_44), inference(backward_demodulation, [], [f269, f774])).
fof(f3894, plain, (~ spl66_50 | ~ spl66_58), inference(avatar_split_clause, [], [f3889, f832, f798])).
fof(f3889, plain, (~ (e1 = op(e0, e3)) | ~ spl66_58), inference(backward_demodulation, [], [f263, f834])).
fof(f3893, plain, (~ spl66_10 | ~ spl66_58), inference(avatar_split_clause, [], [f3888, f832, f628])).
fof(f3888, plain, (~ (e1 = op(e3, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f243, f834])).
fof(f3875, plain, (~ spl66_60 | ~ spl66_21 | ~ spl66_30 | spl66_168), inference(avatar_split_clause, [], [f3846, f1600, f713, f675, f840])).
fof(f1600, plain, (spl66_168 <=> (e3 = op(op(e2, e2), op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl66_168])])).
fof(f3846, plain, (~ (e3 = op(e0, e1)) | (~ spl66_21 | ~ spl66_30 | spl66_168)), inference(forward_demodulation, [], [f3845, f715])).
fof(f3845, plain, (~ (e3 = op(e0, op(e2, e0))) | (~ spl66_21 | spl66_168)), inference(forward_demodulation, [], [f1602, f677])).
fof(f1602, plain, (~ (e3 = op(op(e2, e2), op(e2, op(e2, e2)))) | spl66_168), inference(avatar_component_clause, [], [f1600])).
fof(f3873, plain, (~ spl66_52 | ~ spl66_56), inference(avatar_split_clause, [], [f3872, f823, f806])).
fof(f3872, plain, (~ (e3 = op(e0, e3)) | ~ spl66_56), inference(forward_demodulation, [], [f264, f825])).
fof(f3870, plain, (~ spl66_51 | ~ spl66_56 | spl66_145), inference(avatar_split_clause, [], [f3854, f1485, f823, f802])).
fof(f3854, plain, (~ (e2 = op(e0, e3)) | (~ spl66_56 | spl66_145)), inference(forward_demodulation, [], [f1486, f825])).
fof(f1486, plain, (~ (e2 = op(e0, op(e0, e2))) | spl66_145), inference(avatar_component_clause, [], [f1485])).
fof(f3829, plain, (~ spl66_41 | ~ spl66_47 | ~ spl66_56 | spl66_169), inference(avatar_contradiction_clause, [], [f3828])).
fof(f3828, plain, ($false | (~ spl66_41 | ~ spl66_47 | ~ spl66_56 | spl66_169)), inference(subsumption_resolution, [], [f3827, f825])).
fof(f3827, plain, (~ (e3 = op(e0, e2)) | (~ spl66_41 | ~ spl66_47 | spl66_169)), inference(forward_demodulation, [], [f3826, f787])).
fof(f3826, plain, (~ (e3 = op(e0, op(e1, e0))) | (~ spl66_41 | spl66_169)), inference(forward_demodulation, [], [f1607, f762])).
fof(f1607, plain, (~ (e3 = op(op(e1, e1), op(e1, op(e1, e1)))) | spl66_169), inference(avatar_component_clause, [], [f1605])).
fof(f1605, plain, (spl66_169 <=> (e3 = op(op(e1, e1), op(e1, op(e1, e1))))), introduced(avatar_definition, [new_symbols(naming, [spl66_169])])).
fof(f3801, plain, (~ spl66_19 | ~ spl66_21 | ~ spl66_138), inference(avatar_contradiction_clause, [], [f3800])).
fof(f3800, plain, ($false | (~ spl66_19 | ~ spl66_21 | ~ spl66_138)), inference(subsumption_resolution, [], [f3799, f285])).
fof(f3799, plain, ((e0 = e3) | (~ spl66_19 | ~ spl66_21 | ~ spl66_138)), inference(forward_demodulation, [], [f3797, f677])).
fof(f3797, plain, ((e3 = op(e2, e2)) | (~ spl66_19 | ~ spl66_138)), inference(backward_demodulation, [], [f1457, f668])).
fof(f3791, plain, (~ spl66_21 | spl66_29 | ~ spl66_159), inference(avatar_contradiction_clause, [], [f3790])).
fof(f3790, plain, ($false | (~ spl66_21 | spl66_29 | ~ spl66_159)), inference(subsumption_resolution, [], [f3787, f710])).
fof(f710, plain, (~ (e0 = op(e2, e0)) | spl66_29), inference(avatar_component_clause, [], [f709])).
fof(f3787, plain, ((e0 = op(e2, e0)) | (~ spl66_21 | ~ spl66_159)), inference(backward_demodulation, [], [f1554, f677])).
fof(f1554, plain, ((e0 = op(e2, op(e2, e2))) | ~ spl66_159), inference(avatar_component_clause, [], [f1553])).
fof(f3789, plain, (~ spl66_17 | ~ spl66_21), inference(avatar_split_clause, [], [f3783, f675, f658])).
fof(f3783, plain, (~ (e0 = op(e2, e3)) | ~ spl66_21), inference(backward_demodulation, [], [f276, f677])).
fof(f3764, plain, (~ spl66_22 | ~ spl66_38), inference(avatar_split_clause, [], [f3758, f747, f679])).
fof(f3758, plain, (~ (e1 = op(e2, e2)) | ~ spl66_38), inference(backward_demodulation, [], [f250, f749])).
fof(f3752, plain, (~ spl66_41 | ~ spl66_47 | spl66_153), inference(avatar_contradiction_clause, [], [f3751])).
fof(f3751, plain, ($false | (~ spl66_41 | ~ spl66_47 | spl66_153)), inference(subsumption_resolution, [], [f3743, f787])).
fof(f3743, plain, (~ (e2 = op(e1, e0)) | (~ spl66_41 | spl66_153)), inference(backward_demodulation, [], [f1527, f762])).
fof(f1527, plain, (~ (e2 = op(e1, op(e1, e1))) | spl66_153), inference(avatar_component_clause, [], [f1525])).
fof(f3735, plain, (~ spl66_31 | ~ spl66_47), inference(avatar_split_clause, [], [f3728, f785, f717])).
fof(f3728, plain, (~ (e2 = op(e2, e0)) | ~ spl66_47), inference(backward_demodulation, [], [f238, f787])).
fof(f3727, plain, (~ spl66_50 | ~ spl66_56 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f3726])).
fof(f3726, plain, ($false | (~ spl66_50 | ~ spl66_56 | ~ spl66_145)), inference(subsumption_resolution, [], [f3725, f286])).
fof(f3725, plain, ((e1 = e2) | (~ spl66_50 | ~ spl66_56 | ~ spl66_145)), inference(forward_demodulation, [], [f3722, f800])).
fof(f3722, plain, ((e2 = op(e0, e3)) | (~ spl66_56 | ~ spl66_145)), inference(backward_demodulation, [], [f1487, f825])).
fof(f3708, plain, (~ spl66_53 | ~ spl66_61), inference(avatar_split_clause, [], [f3699, f845, f811])).
fof(f3699, plain, (~ (e0 = op(e0, e2)) | ~ spl66_61), inference(backward_demodulation, [], [f260, f847])).
fof(f3707, plain, (~ spl66_13 | ~ spl66_61), inference(avatar_split_clause, [], [f3697, f845, f641])).
fof(f3697, plain, (~ (e0 = op(e3, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f237, f847])).
fof(f3679, plain, (~ spl66_14 | ~ spl66_10), inference(avatar_split_clause, [], [f3678, f628, f645])).
fof(f3678, plain, (~ (e1 = op(e3, e0)) | ~ spl66_10), inference(forward_demodulation, [], [f277, f630])).
fof(f630, plain, ((e1 = op(e3, e1)) | ~ spl66_10), inference(avatar_component_clause, [], [f628])).
fof(f3676, plain, (~ spl66_3 | ~ spl66_25 | spl66_147 | ~ spl66_148), inference(avatar_contradiction_clause, [], [f3675])).
fof(f3675, plain, ($false | (~ spl66_3 | ~ spl66_25 | spl66_147 | ~ spl66_148)), inference(subsumption_resolution, [], [f3670, f694])).
fof(f3670, plain, (~ (e0 = op(e2, e1)) | (~ spl66_3 | spl66_147 | ~ spl66_148)), inference(backward_demodulation, [], [f3592, f600])).
fof(f3592, plain, (~ (e0 = op(op(e3, e3), e1)) | (spl66_147 | ~ spl66_148)), inference(forward_demodulation, [], [f1500, f1503])).
fof(f1503, plain, ((e1 = op(e3, op(e3, e3))) | ~ spl66_148), inference(avatar_component_clause, [], [f1502])).
fof(f3672, plain, (~ spl66_3 | spl66_6 | ~ spl66_148), inference(avatar_contradiction_clause, [], [f3671])).
fof(f3671, plain, ($false | (~ spl66_3 | spl66_6 | ~ spl66_148)), inference(subsumption_resolution, [], [f3666, f612])).
fof(f612, plain, (~ (e1 = op(e3, e2)) | spl66_6), inference(avatar_component_clause, [], [f611])).
fof(f3666, plain, ((e1 = op(e3, e2)) | (~ spl66_3 | ~ spl66_148)), inference(backward_demodulation, [], [f1503, f600])).
fof(f3649, plain, (~ spl66_21 | ~ spl66_25), inference(avatar_split_clause, [], [f3644, f692, f675])).
fof(f3644, plain, (~ (e0 = op(e2, e2)) | ~ spl66_25), inference(backward_demodulation, [], [f274, f694])).
fof(f3642, plain, (~ spl66_33 | ~ spl66_46 | ~ spl66_142), inference(avatar_contradiction_clause, [], [f3641])).
fof(f3641, plain, ($false | (~ spl66_33 | ~ spl66_46 | ~ spl66_142)), inference(subsumption_resolution, [], [f3640, f287])).
fof(f3640, plain, ((e1 = e3) | (~ spl66_33 | ~ spl66_46 | ~ spl66_142)), inference(forward_demodulation, [], [f3638, f783])).
fof(f3638, plain, ((e3 = op(e1, e0)) | (~ spl66_33 | ~ spl66_142)), inference(backward_demodulation, [], [f1474, f728])).
fof(f3611, plain, (~ spl66_14 | ~ spl66_46), inference(avatar_split_clause, [], [f3604, f781, f645])).
fof(f3604, plain, (~ (e1 = op(e3, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f239, f783])).
fof(f3603, plain, (~ spl66_53 | ~ spl66_55), inference(avatar_contradiction_clause, [], [f3602])).
fof(f3602, plain, ($false | (~ spl66_53 | ~ spl66_55)), inference(subsumption_resolution, [], [f3601, f284])).
fof(f3601, plain, ((e0 = e2) | (~ spl66_53 | ~ spl66_55)), inference(backward_demodulation, [], [f821, f813])).
fof(f3580, plain, (~ spl66_60 | ~ spl66_64), inference(avatar_split_clause, [], [f3475, f857, f840])).
fof(f3475, plain, (~ (e3 = op(e0, e1)) | ~ spl66_64), inference(backward_demodulation, [], [f259, f859])).
fof(f3577, plain, (~ spl66_60 | ~ spl66_50 | spl66_146), inference(avatar_split_clause, [], [f3498, f1489, f798, f840])).
fof(f3498, plain, (~ (e3 = op(e0, e1)) | (~ spl66_50 | spl66_146)), inference(backward_demodulation, [], [f1490, f800])).
fof(f1490, plain, (~ (e3 = op(e0, op(e0, e3))) | spl66_146), inference(avatar_component_clause, [], [f1489])).
fof(f3567, plain, (~ spl66_12 | ~ spl66_8), inference(avatar_split_clause, [], [f3566, f619, f636])).
fof(f3566, plain, (~ (e3 = op(e3, e1)) | ~ spl66_8), inference(forward_demodulation, [], [f280, f621])).
fof(f621, plain, ((e3 = op(e3, e2)) | ~ spl66_8), inference(avatar_component_clause, [], [f619])).
fof(f3563, plain, (~ spl66_4 | ~ spl66_8), inference(avatar_split_clause, [], [f3562, f619, f602])).
fof(f3562, plain, (~ (e3 = op(e3, e3)) | ~ spl66_8), inference(forward_demodulation, [], [f282, f621])).
fof(f3541, plain, (~ spl66_21 | ~ spl66_50 | ~ spl66_154 | spl66_158), inference(avatar_contradiction_clause, [], [f3540])).
fof(f3540, plain, ($false | (~ spl66_21 | ~ spl66_50 | ~ spl66_154 | spl66_158)), inference(subsumption_resolution, [], [f3533, f800])).
fof(f3533, plain, (~ (e1 = op(e0, e3)) | (~ spl66_21 | ~ spl66_154 | spl66_158)), inference(backward_demodulation, [], [f3349, f677])).
fof(f3349, plain, (~ (e1 = op(op(e2, e2), e3)) | (~ spl66_154 | spl66_158)), inference(backward_demodulation, [], [f1551, f1531])).
fof(f3502, plain, (~ spl66_18 | ~ spl66_50), inference(avatar_split_clause, [], [f3499, f798, f662])).
fof(f3499, plain, (~ (e1 = op(e2, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f254, f800])).
fof(f3495, plain, (~ spl66_25 | ~ spl66_57), inference(avatar_split_clause, [], [f3491, f828, f692])).
fof(f3491, plain, (~ (e0 = op(e2, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f242, f830])).
fof(f3482, plain, (~ spl66_16 | ~ spl66_64), inference(avatar_split_clause, [], [f3474, f857, f653])).
fof(f3474, plain, (~ (e3 = op(e3, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f237, f859])).
fof(f3470, plain, (spl66_16 | ~ spl66_1 | ~ spl66_134), inference(avatar_split_clause, [], [f3469, f1438, f590, f653])).
fof(f3469, plain, ((e3 = op(e3, e0)) | (~ spl66_1 | ~ spl66_134)), inference(forward_demodulation, [], [f1440, f592])).
fof(f3442, plain, (~ spl66_17 | ~ spl66_1), inference(avatar_split_clause, [], [f3441, f590, f658])).
fof(f3441, plain, (~ (e0 = op(e2, e3)) | ~ spl66_1), inference(forward_demodulation, [], [f258, f592])).
fof(f3440, plain, (~ spl66_19 | ~ spl66_31), inference(avatar_split_clause, [], [f3439, f717, f666])).
fof(f3439, plain, (~ (e2 = op(e2, e3)) | ~ spl66_31), inference(forward_demodulation, [], [f273, f719])).
fof(f3435, plain, (~ spl66_7 | ~ spl66_55), inference(avatar_split_clause, [], [f3434, f819, f615])).
fof(f3434, plain, (~ (e2 = op(e3, e2)) | ~ spl66_55), inference(forward_demodulation, [], [f249, f821])).
fof(f3430, plain, (~ spl66_13 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f3429])).
fof(f3429, plain, ($false | (~ spl66_13 | ~ spl66_16)), inference(subsumption_resolution, [], [f3428, f285])).
fof(f3428, plain, ((e0 = e3) | (~ spl66_13 | ~ spl66_16)), inference(backward_demodulation, [], [f655, f643])).
fof(f3414, plain, (~ spl66_18 | ~ spl66_24 | spl66_150), inference(avatar_contradiction_clause, [], [f3413])).
fof(f3413, plain, ($false | (~ spl66_18 | ~ spl66_24 | spl66_150)), inference(subsumption_resolution, [], [f3408, f664])).
fof(f3408, plain, (~ (e1 = op(e2, e3)) | (~ spl66_24 | spl66_150)), inference(backward_demodulation, [], [f1513, f689])).
fof(f3397, plain, (~ spl66_15 | ~ spl66_31), inference(avatar_split_clause, [], [f3392, f717, f649])).
fof(f3392, plain, (~ (e2 = op(e3, e0)) | ~ spl66_31), inference(backward_demodulation, [], [f240, f719])).
fof(f3381, plain, (~ spl66_12 | ~ spl66_44), inference(avatar_split_clause, [], [f3373, f772, f636])).
fof(f3373, plain, (~ (e3 = op(e3, e1)) | ~ spl66_44), inference(backward_demodulation, [], [f245, f774])).
fof(f3371, plain, (~ spl66_49 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f3370])).
fof(f3370, plain, ($false | (~ spl66_49 | ~ spl66_52)), inference(subsumption_resolution, [], [f3369, f285])).
fof(f3369, plain, ((e0 = e3) | (~ spl66_49 | ~ spl66_52)), inference(backward_demodulation, [], [f808, f796])).
fof(f3362, plain, (~ spl66_15 | ~ spl66_1 | spl66_151), inference(avatar_split_clause, [], [f3361, f1516, f590, f649])).
fof(f3361, plain, (~ (e2 = op(e3, e0)) | (~ spl66_1 | spl66_151)), inference(forward_demodulation, [], [f1518, f592])).
fof(f3356, plain, (spl66_13 | ~ spl66_1 | ~ spl66_157), inference(avatar_split_clause, [], [f3355, f1544, f590, f641])).
fof(f3355, plain, ((e0 = op(e3, e0)) | (~ spl66_1 | ~ spl66_157)), inference(forward_demodulation, [], [f1545, f592])).
fof(f1545, plain, ((e0 = op(e3, op(e3, e3))) | ~ spl66_157), inference(avatar_component_clause, [], [f1544])).
fof(f3351, plain, (spl66_28 | ~ spl66_18 | ~ spl66_138), inference(avatar_split_clause, [], [f3334, f1455, f662, f704])).
fof(f3334, plain, ((e3 = op(e2, e1)) | (~ spl66_18 | ~ spl66_138)), inference(forward_demodulation, [], [f1457, f664])).
fof(f3305, plain, (~ spl66_21 | spl66_31 | ~ spl66_137), inference(avatar_contradiction_clause, [], [f3304])).
fof(f3304, plain, ($false | (~ spl66_21 | spl66_31 | ~ spl66_137)), inference(subsumption_resolution, [], [f3299, f718])).
fof(f718, plain, (~ (e2 = op(e2, e0)) | spl66_31), inference(avatar_component_clause, [], [f717])).
fof(f3299, plain, ((e2 = op(e2, e0)) | (~ spl66_21 | ~ spl66_137)), inference(backward_demodulation, [], [f1453, f677])).
fof(f3294, plain, (~ spl66_11 | ~ spl66_27), inference(avatar_split_clause, [], [f3291, f700, f632])).
fof(f3291, plain, (~ (e2 = op(e3, e1)) | ~ spl66_27), inference(backward_demodulation, [], [f246, f702])).
fof(f3284, plain, (~ spl66_41 | ~ spl66_46 | ~ spl66_153), inference(avatar_contradiction_clause, [], [f3283])).
fof(f3283, plain, ($false | (~ spl66_41 | ~ spl66_46 | ~ spl66_153)), inference(subsumption_resolution, [], [f3282, f286])).
fof(f3282, plain, ((e1 = e2) | (~ spl66_41 | ~ spl66_46 | ~ spl66_153)), inference(forward_demodulation, [], [f3273, f783])).
fof(f3273, plain, ((e2 = op(e1, e0)) | (~ spl66_41 | ~ spl66_153)), inference(backward_demodulation, [], [f1526, f762])).
fof(f3281, plain, (~ spl66_41 | ~ spl66_46 | spl66_140), inference(avatar_contradiction_clause, [], [f3280])).
fof(f3280, plain, ($false | (~ spl66_41 | ~ spl66_46 | spl66_140)), inference(subsumption_resolution, [], [f3272, f783])).
fof(f3272, plain, (~ (e1 = op(e1, e0)) | (~ spl66_41 | spl66_140)), inference(backward_demodulation, [], [f1465, f762])).
fof(f3279, plain, (~ spl66_37 | ~ spl66_41), inference(avatar_split_clause, [], [f3268, f760, f743])).
fof(f3268, plain, (~ (e0 = op(e1, e2)) | ~ spl66_41), inference(backward_demodulation, [], [f268, f762])).
fof(f3258, plain, (~ spl66_137 | ~ spl66_154), inference(avatar_contradiction_clause, [], [f3257])).
fof(f3257, plain, ($false | (~ spl66_137 | ~ spl66_154)), inference(subsumption_resolution, [], [f3256, f288])).
fof(f3256, plain, ((e2 = e3) | (~ spl66_137 | ~ spl66_154)), inference(forward_demodulation, [], [f1531, f1453])).
fof(f3242, plain, (~ spl66_58 | spl66_144), inference(avatar_contradiction_clause, [], [f3241])).
fof(f3241, plain, ($false | (~ spl66_58 | spl66_144)), inference(subsumption_resolution, [], [f3240, f834])).
fof(f3240, plain, (~ (e1 = op(e0, e1)) | (~ spl66_58 | spl66_144)), inference(forward_demodulation, [], [f1482, f834])).
fof(f3225, plain, (spl66_11 | ~ spl66_6 | ~ spl66_133), inference(avatar_split_clause, [], [f3224, f1434, f611, f632])).
fof(f1434, plain, (spl66_133 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_133])])).
fof(f3224, plain, ((e2 = op(e3, e1)) | (~ spl66_6 | ~ spl66_133)), inference(forward_demodulation, [], [f1436, f613])).
fof(f1436, plain, ((e2 = op(e3, op(e3, e2))) | ~ spl66_133), inference(avatar_component_clause, [], [f1434])).
fof(f3216, plain, (~ spl66_6 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f3215])).
fof(f3215, plain, ($false | (~ spl66_6 | ~ spl66_8)), inference(subsumption_resolution, [], [f3214, f287])).
fof(f3214, plain, ((e1 = e3) | (~ spl66_6 | ~ spl66_8)), inference(backward_demodulation, [], [f621, f613])).
fof(f3212, plain, (spl66_3 | ~ spl66_8 | ~ spl66_133), inference(avatar_contradiction_clause, [], [f3211])).
fof(f3211, plain, ($false | (spl66_3 | ~ spl66_8 | ~ spl66_133)), inference(subsumption_resolution, [], [f3209, f599])).
fof(f599, plain, (~ (e2 = op(e3, e3)) | spl66_3), inference(avatar_component_clause, [], [f598])).
fof(f3209, plain, ((e2 = op(e3, e3)) | (~ spl66_8 | ~ spl66_133)), inference(backward_demodulation, [], [f1436, f621])).
fof(f3189, plain, (spl66_22 | ~ spl66_23 | ~ spl66_137 | ~ spl66_158), inference(avatar_split_clause, [], [f3185, f1549, f1451, f683, f679])).
fof(f3185, plain, ((e1 = op(e2, e2)) | (~ spl66_23 | ~ spl66_137 | ~ spl66_158)), inference(backward_demodulation, [], [f3092, f685])).
fof(f3092, plain, ((e1 = op(op(e2, e2), e2)) | (~ spl66_137 | ~ spl66_158)), inference(backward_demodulation, [], [f1550, f1453])).
fof(f1550, plain, ((e1 = op(op(e2, e2), op(e2, op(e2, e2)))) | ~ spl66_158), inference(avatar_component_clause, [], [f1549])).
fof(f3150, plain, (~ spl66_33 | ~ spl66_44 | spl66_165), inference(avatar_split_clause, [], [f3142, f1583, f772, f726])).
fof(f3142, plain, (~ (e0 = op(e1, e3)) | (~ spl66_44 | spl66_165)), inference(backward_demodulation, [], [f1585, f774])).
fof(f3149, plain, (spl66_34 | ~ spl66_44 | ~ spl66_140), inference(avatar_contradiction_clause, [], [f3148])).
fof(f3148, plain, ($false | (spl66_34 | ~ spl66_44 | ~ spl66_140)), inference(subsumption_resolution, [], [f3141, f731])).
fof(f731, plain, (~ (e1 = op(e1, e3)) | spl66_34), inference(avatar_component_clause, [], [f730])).
fof(f3141, plain, ((e1 = op(e1, e3)) | (~ spl66_44 | ~ spl66_140)), inference(backward_demodulation, [], [f1466, f774])).
fof(f3136, plain, (~ spl66_42 | ~ spl66_46), inference(avatar_split_clause, [], [f3129, f781, f764])).
fof(f3129, plain, (~ (e1 = op(e1, e1)) | ~ spl66_46), inference(backward_demodulation, [], [f265, f783])).
fof(f3105, plain, (~ spl66_29 | ~ spl66_61), inference(avatar_split_clause, [], [f3094, f845, f709])).
fof(f3094, plain, (~ (e0 = op(e2, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f236, f847])).
fof(f3069, plain, (spl66_9 | ~ spl66_14 | ~ spl66_131), inference(avatar_split_clause, [], [f3067, f1426, f645, f624])).
fof(f3067, plain, ((e0 = op(e3, e1)) | (~ spl66_14 | ~ spl66_131)), inference(backward_demodulation, [], [f1428, f647])).
fof(f3060, plain, (~ spl66_22 | ~ spl66_28 | ~ spl66_33 | spl66_149), inference(avatar_contradiction_clause, [], [f3059])).
fof(f3059, plain, ($false | (~ spl66_22 | ~ spl66_28 | ~ spl66_33 | spl66_149)), inference(subsumption_resolution, [], [f3057, f728])).
fof(f3057, plain, (~ (e0 = op(e1, e3)) | (~ spl66_22 | ~ spl66_28 | spl66_149)), inference(backward_demodulation, [], [f2730, f706])).
fof(f2730, plain, (~ (e0 = op(e1, op(e2, e1))) | (~ spl66_22 | spl66_149)), inference(backward_demodulation, [], [f1509, f681])).
fof(f1509, plain, (~ (e0 = op(op(e2, e2), op(e2, op(e2, e2)))) | spl66_149), inference(avatar_component_clause, [], [f1507])).
fof(f1507, plain, (spl66_149 <=> (e0 = op(op(e2, e2), op(e2, op(e2, e2))))), introduced(avatar_definition, [new_symbols(naming, [spl66_149])])).
fof(f3037, plain, (~ spl66_10 | ~ spl66_42), inference(avatar_split_clause, [], [f3028, f764, f628])).
fof(f3028, plain, (~ (e1 = op(e3, e1)) | ~ spl66_42), inference(backward_demodulation, [], [f245, f766])).
fof(f3009, plain, (~ spl66_56 | ~ spl66_64), inference(avatar_split_clause, [], [f3002, f857, f823])).
fof(f3002, plain, (~ (e3 = op(e0, e2)) | ~ spl66_64), inference(backward_demodulation, [], [f260, f859])).
fof(f3008, plain, (~ spl66_32 | ~ spl66_64), inference(avatar_split_clause, [], [f3000, f857, f721])).
fof(f3000, plain, (~ (e3 = op(e2, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f236, f859])).
fof(f2998, plain, (~ spl66_140 | ~ spl66_165), inference(avatar_contradiction_clause, [], [f2997])).
fof(f2997, plain, ($false | (~ spl66_140 | ~ spl66_165)), inference(subsumption_resolution, [], [f2996, f283])).
fof(f283, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f2996, plain, ((e0 = e1) | (~ spl66_140 | ~ spl66_165)), inference(forward_demodulation, [], [f1466, f1584])).
fof(f2979, plain, (~ spl66_28 | ~ spl66_22 | spl66_154), inference(avatar_split_clause, [], [f2732, f1530, f679, f704])).
fof(f2732, plain, (~ (e3 = op(e2, e1)) | (~ spl66_22 | spl66_154)), inference(backward_demodulation, [], [f1532, f681])).
fof(f1532, plain, (~ (e3 = op(e2, op(e2, e2))) | spl66_154), inference(avatar_component_clause, [], [f1530])).
fof(f2947, plain, (~ spl66_49 | ~ spl66_50), inference(avatar_contradiction_clause, [], [f2946])).
fof(f2946, plain, ($false | (~ spl66_49 | ~ spl66_50)), inference(subsumption_resolution, [], [f2945, f283])).
fof(f2945, plain, ((e0 = e1) | (~ spl66_49 | ~ spl66_50)), inference(backward_demodulation, [], [f800, f796])).
fof(f2942, plain, (spl66_53 | ~ spl66_63 | ~ spl66_143), inference(avatar_split_clause, [], [f2937, f1477, f853, f811])).
fof(f2937, plain, ((e0 = op(e0, e2)) | (~ spl66_63 | ~ spl66_143)), inference(backward_demodulation, [], [f1479, f855])).
fof(f2927, plain, (spl66_35 | ~ spl66_44 | ~ spl66_153), inference(avatar_split_clause, [], [f2926, f1525, f772, f734])).
fof(f2926, plain, ((e2 = op(e1, e3)) | (~ spl66_44 | ~ spl66_153)), inference(forward_demodulation, [], [f1526, f774])).
fof(f2909, plain, (~ spl66_48 | ~ spl66_44), inference(avatar_split_clause, [], [f2908, f772, f789])).
fof(f2908, plain, (~ (e3 = op(e1, e0)) | ~ spl66_44), inference(forward_demodulation, [], [f265, f774])).
fof(f2907, plain, (~ spl66_48 | ~ spl66_22 | ~ spl66_25 | spl66_168), inference(avatar_split_clause, [], [f2790, f1600, f692, f679, f789])).
fof(f2790, plain, (~ (e3 = op(e1, e0)) | (~ spl66_22 | ~ spl66_25 | spl66_168)), inference(forward_demodulation, [], [f2789, f694])).
fof(f2789, plain, (~ (e3 = op(e1, op(e2, e1))) | (~ spl66_22 | spl66_168)), inference(forward_demodulation, [], [f1602, f681])).
fof(f2905, plain, (~ spl66_36 | ~ spl66_4), inference(avatar_split_clause, [], [f2904, f602, f738])).
fof(f2904, plain, (~ (e3 = op(e1, e3)) | ~ spl66_4), inference(forward_demodulation, [], [f257, f604])).
fof(f2887, plain, (~ spl66_38 | spl66_43 | ~ spl66_141), inference(avatar_contradiction_clause, [], [f2886])).
fof(f2886, plain, ($false | (~ spl66_38 | spl66_43 | ~ spl66_141)), inference(subsumption_resolution, [], [f2884, f769])).
fof(f769, plain, (~ (e2 = op(e1, e1)) | spl66_43), inference(avatar_component_clause, [], [f768])).
fof(f2884, plain, ((e2 = op(e1, e1)) | (~ spl66_38 | ~ spl66_141)), inference(backward_demodulation, [], [f1470, f749])).
fof(f2882, plain, (~ spl66_53 | ~ spl66_56), inference(avatar_contradiction_clause, [], [f2881])).
fof(f2881, plain, ($false | (~ spl66_53 | ~ spl66_56)), inference(subsumption_resolution, [], [f2879, f285])).
fof(f2879, plain, ((e0 = e3) | (~ spl66_53 | ~ spl66_56)), inference(backward_demodulation, [], [f825, f813])).
fof(f2876, plain, (~ spl66_19 | ~ spl66_22 | ~ spl66_138), inference(avatar_contradiction_clause, [], [f2875])).
fof(f2875, plain, ($false | (~ spl66_19 | ~ spl66_22 | ~ spl66_138)), inference(subsumption_resolution, [], [f2874, f287])).
fof(f2874, plain, ((e1 = e3) | (~ spl66_19 | ~ spl66_22 | ~ spl66_138)), inference(forward_demodulation, [], [f2873, f681])).
fof(f2873, plain, ((e3 = op(e2, e2)) | (~ spl66_19 | ~ spl66_138)), inference(forward_demodulation, [], [f1457, f668])).
fof(f2849, plain, (~ spl66_39 | ~ spl66_47), inference(avatar_split_clause, [], [f2848, f785, f751])).
fof(f2848, plain, (~ (e2 = op(e1, e2)) | ~ spl66_47), inference(forward_demodulation, [], [f266, f787])).
fof(f2831, plain, (~ spl66_13 | ~ spl66_14), inference(avatar_contradiction_clause, [], [f2830])).
fof(f2830, plain, ($false | (~ spl66_13 | ~ spl66_14)), inference(subsumption_resolution, [], [f2829, f283])).
fof(f2829, plain, ((e0 = e1) | (~ spl66_13 | ~ spl66_14)), inference(backward_demodulation, [], [f647, f643])).
fof(f2821, plain, (~ spl66_42 | ~ spl66_44), inference(avatar_contradiction_clause, [], [f2820])).
fof(f2820, plain, ($false | (~ spl66_42 | ~ spl66_44)), inference(subsumption_resolution, [], [f2819, f287])).
fof(f2819, plain, ((e1 = e3) | (~ spl66_42 | ~ spl66_44)), inference(backward_demodulation, [], [f774, f766])).
fof(f2787, plain, (spl66_60 | ~ spl66_50 | ~ spl66_146), inference(avatar_split_clause, [], [f2711, f1489, f798, f840])).
fof(f2711, plain, ((e3 = op(e0, e1)) | (~ spl66_50 | ~ spl66_146)), inference(backward_demodulation, [], [f1491, f800])).
fof(f2774, plain, (spl66_47 | ~ spl66_37 | ~ spl66_141), inference(avatar_split_clause, [], [f2773, f1468, f743, f785])).
fof(f2773, plain, ((e2 = op(e1, e0)) | (~ spl66_37 | ~ spl66_141)), inference(forward_demodulation, [], [f1470, f745])).
fof(f2765, plain, (~ spl66_25 | spl66_30 | ~ spl66_136), inference(avatar_contradiction_clause, [], [f2764])).
fof(f2764, plain, ($false | (~ spl66_25 | spl66_30 | ~ spl66_136)), inference(subsumption_resolution, [], [f2763, f714])).
fof(f714, plain, (~ (e1 = op(e2, e0)) | spl66_30), inference(avatar_component_clause, [], [f713])).
fof(f2763, plain, ((e1 = op(e2, e0)) | (~ spl66_25 | ~ spl66_136)), inference(forward_demodulation, [], [f1449, f694])).
fof(f2751, plain, (~ spl66_2 | ~ spl66_12 | spl66_132), inference(avatar_split_clause, [], [f2618, f1430, f636, f594])).
fof(f2618, plain, (~ (e1 = op(e3, e3)) | (~ spl66_12 | spl66_132)), inference(forward_demodulation, [], [f1431, f638])).
fof(f2723, plain, (~ spl66_37 | ~ spl66_40), inference(avatar_contradiction_clause, [], [f2722])).
fof(f2722, plain, ($false | (~ spl66_37 | ~ spl66_40)), inference(subsumption_resolution, [], [f2719, f285])).
fof(f2719, plain, ((e0 = e3) | (~ spl66_37 | ~ spl66_40)), inference(backward_demodulation, [], [f757, f745])).
fof(f2718, plain, (~ spl66_46 | ~ spl66_47), inference(avatar_contradiction_clause, [], [f2717])).
fof(f2717, plain, ($false | (~ spl66_46 | ~ spl66_47)), inference(subsumption_resolution, [], [f2716, f286])).
fof(f2716, plain, ((e1 = e2) | (~ spl66_46 | ~ spl66_47)), inference(backward_demodulation, [], [f787, f783])).
fof(f2714, plain, (~ spl66_34 | ~ spl66_50), inference(avatar_split_clause, [], [f2712, f798, f730])).
fof(f2712, plain, (~ (e1 = op(e1, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f253, f800])).
fof(f2702, plain, (~ spl66_50 | ~ spl66_60 | spl66_144), inference(avatar_split_clause, [], [f2699, f1481, f840, f798])).
fof(f2699, plain, (~ (e1 = op(e0, e3)) | (~ spl66_60 | spl66_144)), inference(backward_demodulation, [], [f1482, f842])).
fof(f2657, plain, (~ spl66_21 | ~ spl66_32 | spl66_154), inference(avatar_contradiction_clause, [], [f2656])).
fof(f2656, plain, ($false | (~ spl66_21 | ~ spl66_32 | spl66_154)), inference(subsumption_resolution, [], [f2652, f723])).
fof(f2652, plain, (~ (e3 = op(e2, e0)) | (~ spl66_21 | spl66_154)), inference(backward_demodulation, [], [f1532, f677])).
fof(f2638, plain, (~ spl66_34 | ~ spl66_40 | ~ spl66_141), inference(avatar_contradiction_clause, [], [f2637])).
fof(f2637, plain, ($false | (~ spl66_34 | ~ spl66_40 | ~ spl66_141)), inference(subsumption_resolution, [], [f2636, f286])).
fof(f2636, plain, ((e1 = e2) | (~ spl66_34 | ~ spl66_40 | ~ spl66_141)), inference(forward_demodulation, [], [f2632, f732])).
fof(f2632, plain, ((e2 = op(e1, e3)) | (~ spl66_40 | ~ spl66_141)), inference(backward_demodulation, [], [f1470, f757])).
fof(f2631, plain, (~ spl66_41 | ~ spl66_43), inference(avatar_contradiction_clause, [], [f2630])).
fof(f2630, plain, ($false | (~ spl66_41 | ~ spl66_43)), inference(subsumption_resolution, [], [f2629, f284])).
fof(f2629, plain, ((e0 = e2) | (~ spl66_41 | ~ spl66_43)), inference(backward_demodulation, [], [f770, f762])).
fof(f2627, plain, (~ spl66_47 | ~ spl66_48), inference(avatar_contradiction_clause, [], [f2626])).
fof(f2626, plain, ($false | (~ spl66_47 | ~ spl66_48)), inference(subsumption_resolution, [], [f2625, f288])).
fof(f2625, plain, ((e2 = e3) | (~ spl66_47 | ~ spl66_48)), inference(backward_demodulation, [], [f791, f787])).
fof(f2612, plain, (~ spl66_59 | ~ spl66_1 | ~ spl66_14 | spl66_163), inference(avatar_split_clause, [], [f2611, f1574, f645, f590, f836])).
fof(f2611, plain, (~ (e2 = op(e0, e1)) | (~ spl66_1 | ~ spl66_14 | spl66_163)), inference(forward_demodulation, [], [f2610, f647])).
fof(f2610, plain, (~ (e2 = op(e0, op(e3, e0))) | (~ spl66_1 | spl66_163)), inference(forward_demodulation, [], [f1576, f592])).
fof(f2601, plain, (~ spl66_61 | spl66_143), inference(avatar_contradiction_clause, [], [f2600])).
fof(f2600, plain, ($false | (~ spl66_61 | spl66_143)), inference(subsumption_resolution, [], [f2599, f847])).
fof(f2599, plain, (~ (e0 = op(e0, e0)) | (~ spl66_61 | spl66_143)), inference(forward_demodulation, [], [f1478, f847])).
fof(f1478, plain, (~ (e0 = op(e0, op(e0, e0))) | spl66_143), inference(avatar_component_clause, [], [f1477])).
fof(f2564, plain, (~ spl66_57 | ~ spl66_61), inference(avatar_split_clause, [], [f2563, f845, f828])).
fof(f2563, plain, (~ (e0 = op(e0, e1)) | ~ spl66_61), inference(forward_demodulation, [], [f259, f847])).
fof(f2562, plain, (~ spl66_59 | ~ spl66_54 | spl66_145), inference(avatar_split_clause, [], [f2218, f1485, f815, f836])).
fof(f2218, plain, (~ (e2 = op(e0, e1)) | (~ spl66_54 | spl66_145)), inference(forward_demodulation, [], [f1486, f817])).
fof(f2546, plain, (~ spl66_7 | spl66_133), inference(avatar_contradiction_clause, [], [f2545])).
fof(f2545, plain, ($false | (~ spl66_7 | spl66_133)), inference(subsumption_resolution, [], [f2544, f617])).
fof(f2544, plain, (~ (e2 = op(e3, e2)) | (~ spl66_7 | spl66_133)), inference(backward_demodulation, [], [f1435, f617])).
fof(f1435, plain, (~ (e2 = op(e3, op(e3, e2))) | spl66_133), inference(avatar_component_clause, [], [f1434])).
fof(f2536, plain, (~ spl66_29 | spl66_135), inference(avatar_contradiction_clause, [], [f2535])).
fof(f2535, plain, ($false | (~ spl66_29 | spl66_135)), inference(subsumption_resolution, [], [f2534, f711])).
fof(f2534, plain, (~ (e0 = op(e2, e0)) | (~ spl66_29 | spl66_135)), inference(backward_demodulation, [], [f1444, f711])).
fof(f2523, plain, (spl66_38 | ~ spl66_43 | ~ spl66_140), inference(avatar_contradiction_clause, [], [f2522])).
fof(f2522, plain, ($false | (spl66_38 | ~ spl66_43 | ~ spl66_140)), inference(subsumption_resolution, [], [f2517, f748])).
fof(f748, plain, (~ (e1 = op(e1, e2)) | spl66_38), inference(avatar_component_clause, [], [f747])).
fof(f2517, plain, ((e1 = op(e1, e2)) | (~ spl66_43 | ~ spl66_140)), inference(backward_demodulation, [], [f1466, f770])).
fof(f2521, plain, (~ spl66_39 | ~ spl66_43), inference(avatar_split_clause, [], [f2515, f768, f751])).
fof(f2515, plain, (~ (e2 = op(e1, e2)) | ~ spl66_43), inference(backward_demodulation, [], [f268, f770])).
fof(f2508, plain, (~ spl66_61 | ~ spl66_63), inference(avatar_contradiction_clause, [], [f2507])).
fof(f2507, plain, ($false | (~ spl66_61 | ~ spl66_63)), inference(subsumption_resolution, [], [f2505, f284])).
fof(f2505, plain, ((e0 = e2) | (~ spl66_61 | ~ spl66_63)), inference(backward_demodulation, [], [f855, f847])).
fof(f2469, plain, (~ spl66_2 | ~ spl66_14), inference(avatar_split_clause, [], [f2468, f645, f594])).
fof(f2468, plain, (~ (e1 = op(e3, e3)) | ~ spl66_14), inference(forward_demodulation, [], [f279, f647])).
fof(f2455, plain, (~ spl66_5 | ~ spl66_11 | ~ spl66_132), inference(avatar_contradiction_clause, [], [f2454])).
fof(f2454, plain, ($false | (~ spl66_5 | ~ spl66_11 | ~ spl66_132)), inference(subsumption_resolution, [], [f2453, f283])).
fof(f2453, plain, ((e0 = e1) | (~ spl66_5 | ~ spl66_11 | ~ spl66_132)), inference(forward_demodulation, [], [f2450, f609])).
fof(f609, plain, ((e0 = op(e3, e2)) | ~ spl66_5), inference(avatar_component_clause, [], [f607])).
fof(f2450, plain, ((e1 = op(e3, e2)) | (~ spl66_11 | ~ spl66_132)), inference(backward_demodulation, [], [f1432, f634])).
fof(f2449, plain, (~ spl66_14 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f2448])).
fof(f2448, plain, ($false | (~ spl66_14 | ~ spl66_16)), inference(subsumption_resolution, [], [f2446, f287])).
fof(f2446, plain, ((e1 = e3) | (~ spl66_14 | ~ spl66_16)), inference(backward_demodulation, [], [f655, f647])).
fof(f2445, plain, (~ spl66_17 | ~ spl66_19), inference(avatar_contradiction_clause, [], [f2444])).
fof(f2444, plain, ($false | (~ spl66_17 | ~ spl66_19)), inference(subsumption_resolution, [], [f2443, f284])).
fof(f2443, plain, ((e0 = e2) | (~ spl66_17 | ~ spl66_19)), inference(forward_demodulation, [], [f668, f660])).
fof(f2421, plain, (~ spl66_31 | ~ spl66_63), inference(avatar_split_clause, [], [f2420, f853, f717])).
fof(f2420, plain, (~ (e2 = op(e2, e0)) | ~ spl66_63), inference(forward_demodulation, [], [f236, f855])).
fof(f2419, plain, (~ spl66_28 | ~ spl66_24), inference(avatar_split_clause, [], [f2418, f687, f704])).
fof(f2418, plain, (~ (e3 = op(e2, e1)) | ~ spl66_24), inference(forward_demodulation, [], [f274, f689])).
fof(f2415, plain, (~ spl66_28 | ~ spl66_54 | ~ spl66_63 | spl66_170), inference(avatar_split_clause, [], [f2375, f1611, f853, f815, f704])).
fof(f2375, plain, (~ (e3 = op(e2, e1)) | (~ spl66_54 | ~ spl66_63 | spl66_170)), inference(forward_demodulation, [], [f2374, f817])).
fof(f2374, plain, (~ (e3 = op(e2, op(e0, e2))) | (~ spl66_63 | spl66_170)), inference(forward_demodulation, [], [f1613, f855])).
fof(f2413, plain, (~ spl66_14 | ~ spl66_17 | ~ spl66_24 | spl66_158), inference(avatar_split_clause, [], [f2378, f1549, f687, f658, f645])).
fof(f2378, plain, (~ (e1 = op(e3, e0)) | (~ spl66_17 | ~ spl66_24 | spl66_158)), inference(forward_demodulation, [], [f2377, f660])).
fof(f2377, plain, (~ (e1 = op(e3, op(e2, e3))) | (~ spl66_24 | spl66_158)), inference(forward_demodulation, [], [f1551, f689])).
fof(f2404, plain, (~ spl66_2 | ~ spl66_34), inference(avatar_split_clause, [], [f2403, f730, f594])).
fof(f2403, plain, (~ (e1 = op(e3, e3)) | ~ spl66_34), inference(forward_demodulation, [], [f257, f732])).
fof(f2388, plain, (~ spl66_17 | ~ spl66_30 | ~ spl66_138), inference(avatar_contradiction_clause, [], [f2387])).
fof(f2387, plain, ($false | (~ spl66_17 | ~ spl66_30 | ~ spl66_138)), inference(subsumption_resolution, [], [f2386, f287])).
fof(f2386, plain, ((e1 = e3) | (~ spl66_17 | ~ spl66_30 | ~ spl66_138)), inference(forward_demodulation, [], [f2385, f715])).
fof(f2385, plain, ((e3 = op(e2, e0)) | (~ spl66_17 | ~ spl66_138)), inference(forward_demodulation, [], [f1457, f660])).
fof(f2373, plain, (~ spl66_54 | ~ spl66_63 | spl66_167), inference(avatar_contradiction_clause, [], [f2372])).
fof(f2372, plain, ($false | (~ spl66_54 | ~ spl66_63 | spl66_167)), inference(subsumption_resolution, [], [f2371, f817])).
fof(f2371, plain, (~ (e1 = op(e0, e2)) | (~ spl66_63 | spl66_167)), inference(forward_demodulation, [], [f1595, f855])).
fof(f1595, plain, (~ (e1 = op(e0, op(e0, e0))) | spl66_167), inference(avatar_component_clause, [], [f1593])).
fof(f2346, plain, (spl66_15 | ~ spl66_5 | ~ spl66_133), inference(avatar_split_clause, [], [f2316, f1434, f607, f649])).
fof(f2316, plain, ((e2 = op(e3, e0)) | (~ spl66_5 | ~ spl66_133)), inference(backward_demodulation, [], [f1436, f609])).
fof(f2345, plain, (~ spl66_3 | ~ spl66_5 | ~ spl66_30 | spl66_156), inference(avatar_contradiction_clause, [], [f2344])).
fof(f2344, plain, ($false | (~ spl66_3 | ~ spl66_5 | ~ spl66_30 | spl66_156)), inference(subsumption_resolution, [], [f2343, f715])).
fof(f2343, plain, (~ (e1 = op(e2, e0)) | (~ spl66_3 | ~ spl66_5 | spl66_156)), inference(forward_demodulation, [], [f2342, f609])).
fof(f2342, plain, (~ (e1 = op(e2, op(e3, e2))) | (~ spl66_3 | spl66_156)), inference(forward_demodulation, [], [f1542, f600])).
fof(f1542, plain, (~ (e1 = op(op(e3, e3), op(e3, op(e3, e3)))) | spl66_156), inference(avatar_component_clause, [], [f1540])).
fof(f1540, plain, (spl66_156 <=> (e1 = op(op(e3, e3), op(e3, op(e3, e3))))), introduced(avatar_definition, [new_symbols(naming, [spl66_156])])).
fof(f2339, plain, (~ spl66_15 | ~ spl66_3), inference(avatar_split_clause, [], [f2338, f598, f649])).
fof(f2338, plain, (~ (e2 = op(e3, e0)) | ~ spl66_3), inference(forward_demodulation, [], [f279, f600])).
fof(f2324, plain, (~ spl66_3 | ~ spl66_4), inference(avatar_contradiction_clause, [], [f2323])).
fof(f2323, plain, ($false | (~ spl66_3 | ~ spl66_4)), inference(subsumption_resolution, [], [f2322, f288])).
fof(f2322, plain, ((e2 = e3) | (~ spl66_3 | ~ spl66_4)), inference(backward_demodulation, [], [f604, f600])).
fof(f2311, plain, (~ spl66_49 | ~ spl66_63 | ~ spl66_146), inference(avatar_contradiction_clause, [], [f2310])).
fof(f2310, plain, ($false | (~ spl66_49 | ~ spl66_63 | ~ spl66_146)), inference(subsumption_resolution, [], [f2309, f288])).
fof(f2309, plain, ((e2 = e3) | (~ spl66_49 | ~ spl66_63 | ~ spl66_146)), inference(forward_demodulation, [], [f2303, f855])).
fof(f2303, plain, ((op(e0, e0) = e3) | (~ spl66_49 | ~ spl66_146)), inference(backward_demodulation, [], [f1491, f796])).
fof(f2298, plain, (~ spl66_63 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f2297])).
fof(f2297, plain, ($false | (~ spl66_63 | ~ spl66_64)), inference(subsumption_resolution, [], [f2296, f288])).
fof(f2296, plain, ((e2 = e3) | (~ spl66_63 | ~ spl66_64)), inference(backward_demodulation, [], [f859, f855])).
fof(f2295, plain, (~ spl66_22 | ~ spl66_27 | spl66_136), inference(avatar_split_clause, [], [f2294, f1447, f700, f679])).
fof(f2294, plain, (~ (e1 = op(e2, e2)) | (~ spl66_27 | spl66_136)), inference(forward_demodulation, [], [f1448, f702])).
fof(f1448, plain, (~ (e1 = op(e2, op(e2, e1))) | spl66_136), inference(avatar_component_clause, [], [f1447])).
fof(f2293, plain, (~ spl66_17 | ~ spl66_24 | ~ spl66_137), inference(avatar_contradiction_clause, [], [f2292])).
fof(f2292, plain, ($false | (~ spl66_17 | ~ spl66_24 | ~ spl66_137)), inference(subsumption_resolution, [], [f2291, f284])).
fof(f2291, plain, ((e0 = e2) | (~ spl66_17 | ~ spl66_24 | ~ spl66_137)), inference(forward_demodulation, [], [f2290, f660])).
fof(f2290, plain, ((e2 = op(e2, e3)) | (~ spl66_24 | ~ spl66_137)), inference(forward_demodulation, [], [f1453, f689])).
fof(f2289, plain, (spl66_1 | ~ spl66_4 | ~ spl66_157), inference(avatar_contradiction_clause, [], [f2288])).
fof(f2288, plain, ($false | (spl66_1 | ~ spl66_4 | ~ spl66_157)), inference(subsumption_resolution, [], [f2287, f591])).
fof(f591, plain, (~ (e0 = op(e3, e3)) | spl66_1), inference(avatar_component_clause, [], [f590])).
fof(f2287, plain, ((e0 = op(e3, e3)) | (~ spl66_4 | ~ spl66_157)), inference(forward_demodulation, [], [f1545, f604])).
fof(f2282, plain, (spl66_51 | ~ spl66_64 | ~ spl66_161), inference(avatar_split_clause, [], [f2217, f1563, f857, f802])).
fof(f2217, plain, ((e2 = op(e0, e3)) | (~ spl66_64 | ~ spl66_161)), inference(forward_demodulation, [], [f1564, f859])).
fof(f2271, plain, (~ spl66_22 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f2270])).
fof(f2270, plain, ($false | (~ spl66_22 | ~ spl66_24)), inference(subsumption_resolution, [], [f2269, f287])).
fof(f2269, plain, ((e1 = e3) | (~ spl66_22 | ~ spl66_24)), inference(backward_demodulation, [], [f689, f681])).
fof(f2265, plain, (~ spl66_51 | ~ spl66_54 | ~ spl66_146), inference(avatar_contradiction_clause, [], [f2264])).
fof(f2264, plain, ($false | (~ spl66_51 | ~ spl66_54 | ~ spl66_146)), inference(subsumption_resolution, [], [f2263, f287])).
fof(f2263, plain, ((e1 = e3) | (~ spl66_51 | ~ spl66_54 | ~ spl66_146)), inference(forward_demodulation, [], [f2262, f817])).
fof(f2262, plain, ((e3 = op(e0, e2)) | (~ spl66_51 | ~ spl66_146)), inference(forward_demodulation, [], [f1491, f804])).
fof(f2256, plain, (spl66_22 | ~ spl66_27 | ~ spl66_136), inference(avatar_split_clause, [], [f2234, f1447, f700, f679])).
fof(f2234, plain, ((e1 = op(e2, e2)) | (~ spl66_27 | ~ spl66_136)), inference(backward_demodulation, [], [f1449, f702])).
fof(f2255, plain, (~ spl66_13 | spl66_131), inference(avatar_contradiction_clause, [], [f2254])).
fof(f2254, plain, ($false | (~ spl66_13 | spl66_131)), inference(subsumption_resolution, [], [f2252, f643])).
fof(f2252, plain, (~ (e0 = op(e3, e0)) | (~ spl66_13 | spl66_131)), inference(backward_demodulation, [], [f1427, f643])).
fof(f1427, plain, (~ (e0 = op(e3, op(e3, e0))) | spl66_131), inference(avatar_component_clause, [], [f1426])).
fof(f2228, plain, (~ spl66_15 | ~ spl66_47), inference(avatar_split_clause, [], [f2226, f785, f649])).
fof(f2226, plain, (~ (e2 = op(e3, e0)) | ~ spl66_47), inference(backward_demodulation, [], [f239, f787])).
fof(f2227, plain, (spl66_37 | ~ spl66_47 | ~ spl66_139), inference(avatar_split_clause, [], [f2225, f1460, f785, f743])).
fof(f2225, plain, ((e0 = op(e1, e2)) | (~ spl66_47 | ~ spl66_139)), inference(backward_demodulation, [], [f1462, f787])).
fof(f2222, plain, (~ spl66_57 | spl66_62 | ~ spl66_144), inference(avatar_contradiction_clause, [], [f2221])).
fof(f2221, plain, ($false | (~ spl66_57 | spl66_62 | ~ spl66_144)), inference(subsumption_resolution, [], [f2220, f850])).
fof(f850, plain, (~ (op(e0, e0) = e1) | spl66_62), inference(avatar_component_clause, [], [f849])).
fof(f2220, plain, ((op(e0, e0) = e1) | (~ spl66_57 | ~ spl66_144)), inference(forward_demodulation, [], [f1483, f830])).
fof(f2177, plain, (spl66_3 | ~ spl66_4 | ~ spl66_151), inference(avatar_split_clause, [], [f2176, f1516, f602, f598])).
fof(f2176, plain, ((e2 = op(e3, e3)) | (~ spl66_4 | ~ spl66_151)), inference(forward_demodulation, [], [f1517, f604])).
fof(f2175, plain, (~ spl66_1 | ~ spl66_4), inference(avatar_contradiction_clause, [], [f2174])).
fof(f2174, plain, ($false | (~ spl66_1 | ~ spl66_4)), inference(subsumption_resolution, [], [f2173, f285])).
fof(f2173, plain, ((e0 = e3) | (~ spl66_1 | ~ spl66_4)), inference(forward_demodulation, [], [f604, f592])).
fof(f2169, plain, (~ spl66_19 | ~ spl66_20), inference(avatar_contradiction_clause, [], [f2168])).
fof(f2168, plain, ($false | (~ spl66_19 | ~ spl66_20)), inference(subsumption_resolution, [], [f2167, f288])).
fof(f2167, plain, ((e2 = e3) | (~ spl66_19 | ~ spl66_20)), inference(forward_demodulation, [], [f672, f668])).
fof(f2162, plain, (~ spl66_43 | ~ spl66_44), inference(avatar_contradiction_clause, [], [f2161])).
fof(f2161, plain, ($false | (~ spl66_43 | ~ spl66_44)), inference(subsumption_resolution, [], [f2160, f288])).
fof(f2160, plain, ((e2 = e3) | (~ spl66_43 | ~ spl66_44)), inference(backward_demodulation, [], [f774, f770])).
fof(f2147, plain, (~ spl66_6 | ~ spl66_51 | ~ spl66_64 | spl66_160), inference(avatar_split_clause, [], [f2144, f1559, f857, f802, f611])).
fof(f2144, plain, (~ (e1 = op(e3, e2)) | (~ spl66_51 | ~ spl66_64 | spl66_160)), inference(backward_demodulation, [], [f2127, f804])).
fof(f2127, plain, (~ (e1 = op(e3, op(e0, e3))) | (~ spl66_64 | spl66_160)), inference(forward_demodulation, [], [f1561, f859])).
fof(f2141, plain, (~ spl66_50 | ~ spl66_54), inference(avatar_split_clause, [], [f2136, f815, f798])).
fof(f2136, plain, (~ (e1 = op(e0, e3)) | ~ spl66_54), inference(backward_demodulation, [], [f264, f817])).
fof(f2123, plain, (~ spl66_54 | ~ spl66_1 | ~ spl66_15 | spl66_156), inference(avatar_split_clause, [], [f2096, f1540, f649, f590, f815])).
fof(f2096, plain, (~ (e1 = op(e0, e2)) | (~ spl66_1 | ~ spl66_15 | spl66_156)), inference(backward_demodulation, [], [f2061, f651])).
fof(f2061, plain, (~ (e1 = op(e0, op(e3, e0))) | (~ spl66_1 | spl66_156)), inference(forward_demodulation, [], [f1542, f592])).
fof(f2118, plain, (~ spl66_51 | ~ spl66_64 | spl66_161), inference(avatar_split_clause, [], [f2043, f1563, f857, f802])).
fof(f2043, plain, (~ (e2 = op(e0, e3)) | (~ spl66_64 | spl66_161)), inference(forward_demodulation, [], [f1565, f859])).
fof(f2109, plain, (~ spl66_26 | ~ spl66_30), inference(avatar_split_clause, [], [f2073, f713, f696])).
fof(f2073, plain, (~ (e1 = op(e2, e1)) | ~ spl66_30), inference(backward_demodulation, [], [f271, f715])).
fof(f2107, plain, (~ spl66_6 | ~ spl66_10), inference(avatar_split_clause, [], [f2106, f628, f611])).
fof(f2106, plain, (~ (e1 = op(e3, e2)) | ~ spl66_10), inference(forward_demodulation, [], [f280, f630])).
fof(f2102, plain, (~ spl66_10 | ~ spl66_12), inference(avatar_contradiction_clause, [], [f2101])).
fof(f2101, plain, ($false | (~ spl66_10 | ~ spl66_12)), inference(subsumption_resolution, [], [f2100, f287])).
fof(f2100, plain, ((e1 = e3) | (~ spl66_10 | ~ spl66_12)), inference(forward_demodulation, [], [f638, f630])).
fof(f2099, plain, (spl66_5 | ~ spl66_15 | ~ spl66_131), inference(avatar_contradiction_clause, [], [f2098])).
fof(f2098, plain, ($false | (spl66_5 | ~ spl66_15 | ~ spl66_131)), inference(subsumption_resolution, [], [f2097, f608])).
fof(f608, plain, (~ (e0 = op(e3, e2)) | spl66_5), inference(avatar_component_clause, [], [f607])).
fof(f2097, plain, ((e0 = op(e3, e2)) | (~ spl66_15 | ~ spl66_131)), inference(backward_demodulation, [], [f1428, f651])).
fof(f2093, plain, (~ spl66_21 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f2092])).
fof(f2092, plain, ($false | (~ spl66_21 | ~ spl66_24)), inference(subsumption_resolution, [], [f2091, f285])).
fof(f2091, plain, ((e0 = e3) | (~ spl66_21 | ~ spl66_24)), inference(forward_demodulation, [], [f689, f677])).
fof(f2070, plain, (~ spl66_39 | ~ spl66_40), inference(avatar_contradiction_clause, [], [f2069])).
fof(f2069, plain, ($false | (~ spl66_39 | ~ spl66_40)), inference(subsumption_resolution, [], [f2068, f288])).
fof(f2068, plain, ((e2 = e3) | (~ spl66_39 | ~ spl66_40)), inference(backward_demodulation, [], [f757, f753])).
fof(f2067, plain, (~ spl66_50 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f2066])).
fof(f2066, plain, ($false | (~ spl66_50 | ~ spl66_52)), inference(subsumption_resolution, [], [f2065, f287])).
fof(f2065, plain, ((e1 = e3) | (~ spl66_50 | ~ spl66_52)), inference(backward_demodulation, [], [f808, f800])).
fof(f2063, plain, (spl66_15 | ~ spl66_1 | ~ spl66_151), inference(avatar_split_clause, [], [f2062, f1516, f590, f649])).
fof(f2062, plain, ((e2 = op(e3, e0)) | (~ spl66_1 | ~ spl66_151)), inference(forward_demodulation, [], [f1517, f592])).
fof(f2051, plain, (spl66_49 | ~ spl66_64 | ~ spl66_143), inference(avatar_contradiction_clause, [], [f2050])).
fof(f2050, plain, ($false | (spl66_49 | ~ spl66_64 | ~ spl66_143)), inference(subsumption_resolution, [], [f2049, f795])).
fof(f795, plain, (~ (e0 = op(e0, e3)) | spl66_49), inference(avatar_component_clause, [], [f794])).
fof(f2049, plain, ((e0 = op(e0, e3)) | (~ spl66_64 | ~ spl66_143)), inference(forward_demodulation, [], [f1479, f859])).
fof(f2045, plain, (~ spl66_50 | ~ spl66_64 | spl66_167), inference(avatar_split_clause, [], [f2044, f1593, f857, f798])).
fof(f2044, plain, (~ (e1 = op(e0, e3)) | (~ spl66_64 | spl66_167)), inference(forward_demodulation, [], [f1595, f859])).
fof(f2014, plain, (~ spl66_10 | ~ spl66_11), inference(avatar_contradiction_clause, [], [f2013])).
fof(f2013, plain, ($false | (~ spl66_10 | ~ spl66_11)), inference(subsumption_resolution, [], [f2012, f286])).
fof(f2012, plain, ((e1 = e2) | (~ spl66_10 | ~ spl66_11)), inference(backward_demodulation, [], [f634, f630])).
fof(f2007, plain, (~ spl66_34 | ~ spl66_35), inference(avatar_contradiction_clause, [], [f2006])).
fof(f2006, plain, ($false | (~ spl66_34 | ~ spl66_35)), inference(subsumption_resolution, [], [f2005, f286])).
fof(f2005, plain, ((e1 = e2) | (~ spl66_34 | ~ spl66_35)), inference(backward_demodulation, [], [f736, f732])).
fof(f1999, plain, (~ spl66_62 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f1998])).
fof(f1998, plain, ($false | (~ spl66_62 | ~ spl66_64)), inference(subsumption_resolution, [], [f1997, f287])).
fof(f1997, plain, ((e1 = e3) | (~ spl66_62 | ~ spl66_64)), inference(forward_demodulation, [], [f859, f851])).
fof(f1973, plain, (~ spl66_21 | ~ spl66_22), inference(avatar_contradiction_clause, [], [f1972])).
fof(f1972, plain, ($false | (~ spl66_21 | ~ spl66_22)), inference(subsumption_resolution, [], [f1971, f283])).
fof(f1971, plain, ((e0 = e1) | (~ spl66_21 | ~ spl66_22)), inference(backward_demodulation, [], [f681, f677])).
fof(f1970, plain, (~ spl66_22 | ~ spl66_23), inference(avatar_contradiction_clause, [], [f1969])).
fof(f1969, plain, ($false | (~ spl66_22 | ~ spl66_23)), inference(subsumption_resolution, [], [f1968, f286])).
fof(f1968, plain, ((e1 = e2) | (~ spl66_22 | ~ spl66_23)), inference(backward_demodulation, [], [f685, f681])).
fof(f1963, plain, (spl66_21 | ~ spl66_31 | ~ spl66_135), inference(avatar_split_clause, [], [f1961, f1443, f717, f675])).
fof(f1961, plain, ((e0 = op(e2, e2)) | (~ spl66_31 | ~ spl66_135)), inference(backward_demodulation, [], [f1445, f719])).
fof(f1933, plain, (~ spl66_57 | ~ spl66_60), inference(avatar_contradiction_clause, [], [f1932])).
fof(f1932, plain, ($false | (~ spl66_57 | ~ spl66_60)), inference(subsumption_resolution, [], [f1931, f285])).
fof(f1931, plain, ((e0 = e3) | (~ spl66_57 | ~ spl66_60)), inference(forward_demodulation, [], [f842, f830])).
fof(f1900, plain, (~ spl66_22 | ~ spl66_6), inference(avatar_split_clause, [], [f1899, f611, f679])).
fof(f1899, plain, (~ (e1 = op(e2, e2)) | ~ spl66_6), inference(forward_demodulation, [], [f252, f613])).
fof(f1898, plain, (~ spl66_22 | ~ spl66_18), inference(avatar_split_clause, [], [f1897, f662, f679])).
fof(f1897, plain, (~ (e1 = op(e2, e2)) | ~ spl66_18), inference(forward_demodulation, [], [f276, f664])).
fof(f1874, plain, (~ spl66_23 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f1873])).
fof(f1873, plain, ($false | (~ spl66_23 | ~ spl66_24)), inference(subsumption_resolution, [], [f1872, f288])).
fof(f1872, plain, ((e2 = e3) | (~ spl66_23 | ~ spl66_24)), inference(backward_demodulation, [], [f689, f685])).
fof(f1871, plain, (~ spl66_26 | ~ spl66_27), inference(avatar_contradiction_clause, [], [f1870])).
fof(f1870, plain, ($false | (~ spl66_26 | ~ spl66_27)), inference(subsumption_resolution, [], [f1869, f286])).
fof(f1869, plain, ((e1 = e2) | (~ spl66_26 | ~ spl66_27)), inference(backward_demodulation, [], [f702, f698])).
fof(f1866, plain, (~ spl66_35 | ~ spl66_36), inference(avatar_contradiction_clause, [], [f1865])).
fof(f1865, plain, ($false | (~ spl66_35 | ~ spl66_36)), inference(subsumption_resolution, [], [f1864, f288])).
fof(f1864, plain, ((e2 = e3) | (~ spl66_35 | ~ spl66_36)), inference(backward_demodulation, [], [f740, f736])).
fof(f1863, plain, (~ spl66_37 | ~ spl66_38), inference(avatar_contradiction_clause, [], [f1862])).
fof(f1862, plain, ($false | (~ spl66_37 | ~ spl66_38)), inference(subsumption_resolution, [], [f1861, f283])).
fof(f1861, plain, ((e0 = e1) | (~ spl66_37 | ~ spl66_38)), inference(backward_demodulation, [], [f749, f745])).
fof(f1853, plain, (~ spl66_35 | ~ spl66_44 | spl66_153), inference(avatar_split_clause, [], [f1849, f1525, f772, f734])).
fof(f1849, plain, (~ (e2 = op(e1, e3)) | (~ spl66_44 | spl66_153)), inference(backward_demodulation, [], [f1527, f774])).
fof(f1847, plain, (spl66_33 | ~ spl66_48 | ~ spl66_139), inference(avatar_split_clause, [], [f1845, f1460, f789, f726])).
fof(f1845, plain, ((e0 = op(e1, e3)) | (~ spl66_48 | ~ spl66_139)), inference(backward_demodulation, [], [f1462, f791])).
fof(f1841, plain, (~ spl66_53 | ~ spl66_57), inference(avatar_split_clause, [], [f1837, f828, f811])).
fof(f1837, plain, (~ (e0 = op(e0, e2)) | ~ spl66_57), inference(backward_demodulation, [], [f262, f830])).
fof(f1832, plain, (~ spl66_62 | ~ spl66_63), inference(avatar_contradiction_clause, [], [f1831])).
fof(f1831, plain, ($false | (~ spl66_62 | ~ spl66_63)), inference(subsumption_resolution, [], [f1830, f286])).
fof(f1830, plain, ((e1 = e2) | (~ spl66_62 | ~ spl66_63)), inference(forward_demodulation, [], [f855, f851])).
fof(f1818, plain, (~ spl66_56 | ~ spl66_24), inference(avatar_split_clause, [], [f1817, f687, f823])).
fof(f1817, plain, (~ (e3 = op(e0, e2)) | ~ spl66_24), inference(forward_demodulation, [], [f248, f689])).
fof(f1802, plain, (~ spl66_43 | ~ spl66_27), inference(avatar_split_clause, [], [f1801, f700, f768])).
fof(f1801, plain, (~ (e2 = op(e1, e1)) | ~ spl66_27), inference(forward_demodulation, [], [f244, f702])).
fof(f1783, plain, (~ spl66_5 | ~ spl66_1), inference(avatar_split_clause, [], [f1782, f590, f607])).
fof(f1782, plain, (~ (e0 = op(e3, e2)) | ~ spl66_1), inference(forward_demodulation, [], [f282, f592])).
fof(f1780, plain, (~ spl66_14 | ~ spl66_1 | spl66_148), inference(avatar_split_clause, [], [f1779, f1502, f590, f645])).
fof(f1779, plain, (~ (e1 = op(e3, e0)) | (~ spl66_1 | spl66_148)), inference(forward_demodulation, [], [f1504, f592])).
fof(f1774, plain, (~ spl66_1 | ~ spl66_2), inference(avatar_contradiction_clause, [], [f1773])).
fof(f1773, plain, ($false | (~ spl66_1 | ~ spl66_2)), inference(subsumption_resolution, [], [f1772, f283])).
fof(f1772, plain, ((e0 = e1) | (~ spl66_1 | ~ spl66_2)), inference(backward_demodulation, [], [f596, f592])).
fof(f1771, plain, (~ spl66_2 | ~ spl66_3), inference(avatar_contradiction_clause, [], [f1770])).
fof(f1770, plain, ($false | (~ spl66_2 | ~ spl66_3)), inference(subsumption_resolution, [], [f1769, f286])).
fof(f1769, plain, ((e1 = e2) | (~ spl66_2 | ~ spl66_3)), inference(backward_demodulation, [], [f600, f596])).
fof(f1768, plain, (~ spl66_5 | ~ spl66_3 | spl66_157), inference(avatar_split_clause, [], [f1764, f1544, f598, f607])).
fof(f1764, plain, (~ (e0 = op(e3, e2)) | (~ spl66_3 | spl66_157)), inference(backward_demodulation, [], [f1546, f600])).
fof(f1752, plain, (~ spl66_15 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f1751])).
fof(f1751, plain, ($false | (~ spl66_15 | ~ spl66_16)), inference(subsumption_resolution, [], [f1750, f288])).
fof(f1750, plain, ((e2 = e3) | (~ spl66_15 | ~ spl66_16)), inference(backward_demodulation, [], [f655, f651])).
fof(f1740, plain, (~ spl66_9 | ~ spl66_18 | ~ spl66_24 | spl66_149), inference(avatar_split_clause, [], [f1736, f1507, f687, f662, f624])).
fof(f1736, plain, (~ (e0 = op(e3, e1)) | (~ spl66_18 | ~ spl66_24 | spl66_149)), inference(backward_demodulation, [], [f1728, f664])).
fof(f1728, plain, (~ (e0 = op(e3, op(e2, e3))) | (~ spl66_24 | spl66_149)), inference(backward_demodulation, [], [f1509, f689])).
fof(f1725, plain, (~ spl66_27 | ~ spl66_28), inference(avatar_contradiction_clause, [], [f1724])).
fof(f1724, plain, ($false | (~ spl66_27 | ~ spl66_28)), inference(subsumption_resolution, [], [f1723, f288])).
fof(f1723, plain, ((e2 = e3) | (~ spl66_27 | ~ spl66_28)), inference(backward_demodulation, [], [f706, f702])).
fof(f1702, plain, (~ spl66_38 | ~ spl66_39), inference(avatar_contradiction_clause, [], [f1701])).
fof(f1701, plain, ($false | (~ spl66_38 | ~ spl66_39)), inference(subsumption_resolution, [], [f1700, f286])).
fof(f1700, plain, ((e1 = e2) | (~ spl66_38 | ~ spl66_39)), inference(backward_demodulation, [], [f753, f749])).
fof(f1678, plain, (~ spl66_33 | ~ spl66_45), inference(avatar_split_clause, [], [f1671, f777, f726])).
fof(f1671, plain, (~ (e0 = op(e1, e3)) | ~ spl66_45), inference(backward_demodulation, [], [f267, f779])).
fof(f1676, plain, (~ spl66_41 | ~ spl66_45), inference(avatar_split_clause, [], [f1669, f777, f760])).
fof(f1669, plain, (~ (e0 = op(e1, e1)) | ~ spl66_45), inference(backward_demodulation, [], [f265, f779])).
fof(f1666, plain, (~ spl66_51 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f1665])).
fof(f1665, plain, ($false | (~ spl66_51 | ~ spl66_52)), inference(subsumption_resolution, [], [f1664, f288])).
fof(f1664, plain, ((e2 = e3) | (~ spl66_51 | ~ spl66_52)), inference(backward_demodulation, [], [f808, f804])).
fof(f1655, plain, (~ spl66_21 | ~ spl66_53), inference(avatar_split_clause, [], [f1651, f811, f675])).
fof(f1651, plain, (~ (e0 = op(e2, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f248, f813])).
fof(f1633, plain, (~ spl66_58 | ~ spl66_62), inference(avatar_split_clause, [], [f1623, f849, f832])).
fof(f1623, plain, (~ (e1 = op(e0, e1)) | ~ spl66_62), inference(backward_demodulation, [], [f259, f851])).
fof(f1630, plain, (~ spl66_46 | ~ spl66_62), inference(avatar_split_clause, [], [f1620, f849, f781])).
fof(f1620, plain, (~ (e1 = op(e1, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f235, f851])).
fof(f1616, plain, (~ spl66_170 | ~ spl66_161 | ~ spl66_62), inference(avatar_split_clause, [], [f588, f849, f1563, f1611])).
fof(f588, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((op(e0, e0) = e1) & (e2 = op(e0, op(e0, e0))) & (e3 = op(op(e0, e0), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax29)).
fof(f1615, plain, (~ spl66_169 | ~ spl66_153 | ~ spl66_41), inference(avatar_split_clause, [], [f587, f760, f1525, f1605])).
fof(f587, plain, (~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e0 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, e1), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e0 = op(e1, e1)) & (e2 = op(e1, op(e1, e1))) & (e3 = op(op(e1, e1), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax28)).
fof(f1614, plain, (~ spl66_170 | ~ spl66_167 | ~ spl66_63), inference(avatar_split_clause, [], [f586, f853, f1593, f1611])).
fof(f586, plain, (~ (op(e0, e0) = e2) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (op(e0, e0) = e2) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((op(e0, e0) = e2) & (e1 = op(e0, op(e0, e0))) & (e3 = op(op(e0, e0), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax27)).
fof(f1609, plain, (~ spl66_168 | ~ spl66_150 | ~ spl66_21), inference(avatar_split_clause, [], [f585, f675, f1511, f1600])).
fof(f585, plain, (~ (e0 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e0 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e0 = op(e2, e2)) & (e1 = op(e2, op(e2, e2))) & (e3 = op(op(e2, e2), op(e2, op(e2, e2))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax26)).
fof(f1603, plain, (~ spl66_168 | ~ spl66_159 | ~ spl66_22), inference(avatar_split_clause, [], [f583, f679, f1553, f1600])).
fof(f583, plain, (~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e1 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e1 = op(e2, e2)) & (e0 = op(e2, op(e2, e2))) & (e3 = op(op(e2, e2), op(e2, op(e2, e2))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax24)).
fof(f1598, plain, (~ spl66_166 | ~ spl66_162 | ~ spl66_62), inference(avatar_split_clause, [], [f582, f849, f1569, f1589])).
fof(f582, plain, (~ (op(e0, e0) = e1) | ~ (e3 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (op(e0, e0) = e1) | ~ (e3 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((op(e0, e0) = e1) & (e3 = op(e0, op(e0, e0))) & (e2 = op(op(e0, e0), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax23)).
fof(f1596, plain, (~ spl66_166 | ~ spl66_167 | ~ spl66_64), inference(avatar_split_clause, [], [f580, f857, f1593, f1589])).
fof(f580, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((op(e0, e0) = e3) & (e1 = op(e0, op(e0, e0))) & (e2 = op(op(e0, e0), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax21)).
fof(f1587, plain, (~ spl66_163 | ~ spl66_148 | ~ spl66_1), inference(avatar_split_clause, [], [f579, f590, f1502, f1574])).
fof(f579, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e0 = op(e3, e3)) & (e1 = op(e3, op(e3, e3))) & (e2 = op(op(e3, e3), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax20)).
fof(f1586, plain, (~ spl66_164 | ~ spl66_165 | ~ spl66_44), inference(avatar_split_clause, [], [f578, f772, f1583, f1579])).
fof(f578, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, e1), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, e1), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e3 = op(e1, e1)) & (e0 = op(e1, op(e1, e1))) & (e2 = op(op(e1, e1), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax19)).
fof(f1577, plain, (~ spl66_163 | ~ spl66_157 | ~ spl66_2), inference(avatar_split_clause, [], [f577, f594, f1544, f1574])).
fof(f577, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e1 = op(e3, e3)) & (e0 = op(e3, op(e3, e3))) & (e2 = op(op(e3, e3), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax18)).
fof(f1572, plain, (~ spl66_160 | ~ spl66_162 | ~ spl66_63), inference(avatar_split_clause, [], [f576, f853, f1569, f1559])).
fof(f576, plain, (~ (op(e0, e0) = e2) | ~ (e3 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (op(e0, e0) = e2) | ~ (e3 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((op(e0, e0) = e2) & (e3 = op(e0, op(e0, e0))) & (e1 = op(op(e0, e0), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax17)).
fof(f1567, plain, (~ spl66_158 | ~ spl66_154 | ~ spl66_21), inference(avatar_split_clause, [], [f575, f675, f1530, f1549])).
fof(f575, plain, (~ (e0 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e0 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e0 = op(e2, e2)) & (e3 = op(e2, op(e2, e2))) & (e1 = op(op(e2, e2), op(e2, op(e2, e2))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax16)).
fof(f1566, plain, (~ spl66_160 | ~ spl66_161 | ~ spl66_64), inference(avatar_split_clause, [], [f574, f857, f1563, f1559])).
fof(f574, plain, (~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (op(e0, e0) = e3) | ~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, e0), op(e0, op(e0, e0))))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((op(e0, e0) = e3) & (e2 = op(e0, op(e0, e0))) & (e1 = op(op(e0, e0), op(e0, op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax15)).
fof(f1557, plain, (~ spl66_156 | ~ spl66_151 | ~ spl66_1), inference(avatar_split_clause, [], [f573, f590, f1516, f1540])).
fof(f573, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e0 = op(e3, e3)) & (e2 = op(e3, op(e3, e3))) & (e1 = op(op(e3, e3), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax14)).
fof(f1556, plain, (~ spl66_158 | ~ spl66_159 | ~ spl66_24), inference(avatar_split_clause, [], [f572, f687, f1553, f1549])).
fof(f572, plain, (~ (e3 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e3 = op(e2, e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e3 = op(e2, e2)) & (e0 = op(e2, op(e2, e2))) & (e1 = op(op(e2, e2), op(e2, op(e2, e2))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax13)).
fof(f1547, plain, (~ spl66_156 | ~ spl66_157 | ~ spl66_3), inference(avatar_split_clause, [], [f571, f598, f1544, f1540])).
fof(f571, plain, (~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e2 = op(e3, e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e2 = op(e3, e3)) & (e0 = op(e3, op(e3, e3))) & (e1 = op(op(e3, e3), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax12)).
fof(f1538, plain, (~ spl66_152 | ~ spl66_155 | ~ spl66_43), inference(avatar_split_clause, [], [f570, f768, f1535, f1521])).
fof(f570, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, e1), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, e1), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e2 = op(e1, e1)) & (e3 = op(e1, op(e1, e1))) & (e0 = op(op(e1, e1), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax11)).
fof(f1533, plain, (~ spl66_149 | ~ spl66_154 | ~ spl66_22), inference(avatar_split_clause, [], [f569, f679, f1530, f1507])).
fof(f569, plain, (~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (e1 = op(e2, e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e1 = op(e2, e2)) & (e3 = op(e2, op(e2, e2))) & (e0 = op(op(e2, e2), op(e2, op(e2, e2))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax10)).
fof(f1528, plain, (~ spl66_152 | ~ spl66_153 | ~ spl66_44), inference(avatar_split_clause, [], [f568, f772, f1525, f1521])).
fof(f568, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, e1), op(e1, op(e1, e1))))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, e1), op(e1, op(e1, e1))))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e3 = op(e1, e1)) & (e2 = op(e1, op(e1, e1))) & (e0 = op(op(e1, e1), op(e1, op(e1, e1))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax9)).
fof(f1519, plain, (~ spl66_147 | ~ spl66_151 | ~ spl66_2), inference(avatar_split_clause, [], [f567, f594, f1516, f1498])).
fof(f567, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e1 = op(e3, e3)) & (e2 = op(e3, op(e3, e3))) & (e0 = op(op(e3, e3), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax8)).
fof(f1514, plain, (~ spl66_149 | ~ spl66_150 | ~ spl66_24), inference(avatar_split_clause, [], [f566, f687, f1511, f1507])).
fof(f566, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, e2), op(e2, op(e2, e2))))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e3 = op(e2, e2)) & (e1 = op(e2, op(e2, e2))) & (e0 = op(op(e2, e2), op(e2, op(e2, e2))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax7)).
fof(f1505, plain, (~ spl66_147 | ~ spl66_148 | ~ spl66_3), inference(avatar_split_clause, [], [f565, f598, f1502, f1498])).
fof(f565, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, e3), op(e3, op(e3, e3))))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e2 = op(e3, e3)) & (e1 = op(e3, op(e3, e3))) & (e0 = op(op(e3, e3), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax6)).
fof(f1496, plain, (spl66_130 | spl66_129 | spl66_128 | ~ spl66_49), inference(avatar_split_clause, [], [f553, f794, f1398, f1406, f1414])).
fof(f1414, plain, (spl66_130 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl66_130])])).
fof(f1406, plain, (spl66_129 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl66_129])])).
fof(f1398, plain, (spl66_128 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl66_128])])).
fof(f553, plain, (~ (e0 = op(e0, e3)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f120, plain, (((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3) & ((e3 = op(e3, op(e3, e3))) | (e2 = op(e3, op(e3, e2))) | (e1 = op(e3, op(e3, e1))) | (e0 = op(e3, op(e3, e0)))) & ((e3 = op(e2, op(e2, e3))) | (e2 = op(e2, op(e2, e2))) | (e1 = op(e2, op(e2, e1))) | (e0 = op(e2, op(e2, e0)))) & ((e3 = op(e1, op(e1, e3))) | (e2 = op(e1, op(e1, e2))) | (e1 = op(e1, op(e1, e1))) | (e0 = op(e1, op(e1, e0)))) & ((e3 = op(e0, op(e0, e3))) | (e2 = op(e0, op(e0, e2))) | (e1 = op(e0, op(e0, e1))) | (e0 = op(e0, op(e0, e0)))) & ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f84, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP30), inference(usedef, [], [e84])).
fof(e84, plain, (sP30 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f85, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(usedef, [], [e85])).
fof(e85, plain, (sP31 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f86, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP32), inference(usedef, [], [e86])).
fof(e86, plain, (sP32 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f87, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP33), inference(usedef, [], [e87])).
fof(e87, plain, (sP33 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f88, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP34), inference(usedef, [], [e88])).
fof(e88, plain, (sP34 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f89, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(usedef, [], [e89])).
fof(e89, plain, (sP35 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f90, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP36), inference(usedef, [], [e90])).
fof(e90, plain, (sP36 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f91, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP37), inference(usedef, [], [e91])).
fof(e91, plain, (sP37 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f92, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP38), inference(usedef, [], [e92])).
fof(e92, plain, (sP38 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f93, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(usedef, [], [e93])).
fof(e93, plain, (sP39 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f94, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP40), inference(usedef, [], [e94])).
fof(e94, plain, (sP40 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f95, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP41), inference(usedef, [], [e95])).
fof(e95, plain, (sP41 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f96, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP42), inference(usedef, [], [e96])).
fof(e96, plain, (sP42 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f97, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(usedef, [], [e97])).
fof(e97, plain, (sP43 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f98, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP44), inference(usedef, [], [e98])).
fof(e98, plain, (sP44 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f99, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP45), inference(usedef, [], [e99])).
fof(e99, plain, (sP45 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f100, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP46), inference(usedef, [], [e100])).
fof(e100, plain, (sP46 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f101, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(usedef, [], [e101])).
fof(e101, plain, (sP47 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f102, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP48), inference(usedef, [], [e102])).
fof(e102, plain, (sP48 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f103, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP49), inference(usedef, [], [e103])).
fof(e103, plain, (sP49 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f104, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP50), inference(usedef, [], [e104])).
fof(e104, plain, (sP50 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f105, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(usedef, [], [e105])).
fof(e105, plain, (sP51 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f106, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP52), inference(usedef, [], [e106])).
fof(e106, plain, (sP52 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f107, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP53), inference(usedef, [], [e107])).
fof(e107, plain, (sP53 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f108, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP54), inference(usedef, [], [e108])).
fof(e108, plain, (sP54 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f109, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(usedef, [], [e109])).
fof(e109, plain, (sP55 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f110, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP56), inference(usedef, [], [e110])).
fof(e110, plain, (sP56 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f111, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP57), inference(usedef, [], [e111])).
fof(e111, plain, (sP57 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f112, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP58), inference(usedef, [], [e112])).
fof(e112, plain, (sP58 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f113, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(usedef, [], [e113])).
fof(e113, plain, (sP59 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f114, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP60), inference(usedef, [], [e114])).
fof(e114, plain, (sP60 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f115, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP61), inference(usedef, [], [e115])).
fof(e115, plain, (sP61 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f116, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP62), inference(usedef, [], [e116])).
fof(e116, plain, (sP62 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f117, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP63), inference(usedef, [], [e117])).
fof(e117, plain, (sP63 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f118, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP64), inference(usedef, [], [e118])).
fof(e118, plain, (sP64 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f119, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP65), inference(usedef, [], [e119])).
fof(e119, plain, (sP65 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f5, plain, (((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((e3 = op(e3, op(e3, e3))) | (e2 = op(e3, op(e3, e2))) | (e1 = op(e3, op(e3, e1))) | (e0 = op(e3, op(e3, e0)))) & ((e3 = op(e2, op(e2, e3))) | (e2 = op(e2, op(e2, e2))) | (e1 = op(e2, op(e2, e1))) | (e0 = op(e2, op(e2, e0)))) & ((e3 = op(e1, op(e1, e3))) | (e2 = op(e1, op(e1, e2))) | (e1 = op(e1, op(e1, e1))) | (e0 = op(e1, op(e1, e0)))) & ((e3 = op(e0, op(e0, e3))) | (e2 = op(e0, op(e0, e2))) | (e1 = op(e0, op(e0, e1))) | (e0 = op(e0, op(e0, e0)))) & ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax5)).
fof(f1495, plain, (spl66_130 | spl66_129 | spl66_128 | ~ spl66_34), inference(avatar_split_clause, [], [f554, f730, f1398, f1406, f1414])).
fof(f554, plain, (~ (e1 = op(e1, e3)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1493, plain, (spl66_130 | spl66_129 | spl66_128 | ~ spl66_4), inference(avatar_split_clause, [], [f556, f602, f1398, f1406, f1414])).
fof(f556, plain, (~ (e3 = op(e3, e3)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1492, plain, (spl66_143 | spl66_144 | spl66_145 | spl66_146), inference(avatar_split_clause, [], [f557, f1489, f1485, f1481, f1477])).
fof(f557, plain, ((e3 = op(e0, op(e0, e3))) | (e2 = op(e0, op(e0, e2))) | (e1 = op(e0, op(e0, e1))) | (e0 = op(e0, op(e0, e0)))), inference(cnf_transformation, [], [f120])).
fof(f1475, plain, (spl66_139 | spl66_140 | spl66_141 | spl66_142), inference(avatar_split_clause, [], [f558, f1472, f1468, f1464, f1460])).
fof(f558, plain, ((e3 = op(e1, op(e1, e3))) | (e2 = op(e1, op(e1, e2))) | (e1 = op(e1, op(e1, e1))) | (e0 = op(e1, op(e1, e0)))), inference(cnf_transformation, [], [f120])).
fof(f1458, plain, (spl66_135 | spl66_136 | spl66_137 | spl66_138), inference(avatar_split_clause, [], [f559, f1455, f1451, f1447, f1443])).
fof(f559, plain, ((e3 = op(e2, op(e2, e3))) | (e2 = op(e2, op(e2, e2))) | (e1 = op(e2, op(e2, e1))) | (e0 = op(e2, op(e2, e0)))), inference(cnf_transformation, [], [f120])).
fof(f1441, plain, (spl66_131 | spl66_132 | spl66_133 | spl66_134), inference(avatar_split_clause, [], [f560, f1438, f1434, f1430, f1426])).
fof(f560, plain, ((e3 = op(e3, op(e3, e3))) | (e2 = op(e3, op(e3, e2))) | (e1 = op(e3, op(e3, e1))) | (e0 = op(e3, op(e3, e0)))), inference(cnf_transformation, [], [f120])).
fof(f1424, plain, (spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | spl66_79 | spl66_78 | spl66_77 | spl66_76 | spl66_75 | spl66_74 | spl66_73 | spl66_72 | spl66_71 | spl66_70 | spl66_69 | spl66_68 | spl66_67 | spl66_66 | spl66_65 | spl66_4), inference(avatar_split_clause, [], [f561, f602, f894, f902, f910, f918, f926, f934, f942, f950, f958, f966, f974, f982, f990, f998, f1006, f1014, f1022, f1030, f1038, f1046, f1054, f1062, f1070, f1078, f1086, f1094, f1102, f1110, f1118, f1126, f1134, f1142, f1150, f1158, f1166, f1174, f1182, f1190, f1198, f1206, f1214, f1222, f1230, f1238, f1246, f1254, f1262, f1270, f1278, f1286, f1294, f1302, f1310, f1318, f1326, f1334, f1342, f1350, f1358, f1366, f1374, f1382, f1390])).
fof(f1390, plain, (spl66_127 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl66_127])])).
fof(f1382, plain, (spl66_126 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl66_126])])).
fof(f1374, plain, (spl66_125 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl66_125])])).
fof(f1366, plain, (spl66_124 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl66_124])])).
fof(f1358, plain, (spl66_123 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl66_123])])).
fof(f1350, plain, (spl66_122 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl66_122])])).
fof(f1342, plain, (spl66_121 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl66_121])])).
fof(f1334, plain, (spl66_120 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl66_120])])).
fof(f1326, plain, (spl66_119 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl66_119])])).
fof(f1318, plain, (spl66_118 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl66_118])])).
fof(f1310, plain, (spl66_117 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl66_117])])).
fof(f1302, plain, (spl66_116 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl66_116])])).
fof(f1294, plain, (spl66_115 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl66_115])])).
fof(f1286, plain, (spl66_114 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl66_114])])).
fof(f1278, plain, (spl66_113 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl66_113])])).
fof(f1270, plain, (spl66_112 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl66_112])])).
fof(f1262, plain, (spl66_111 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl66_111])])).
fof(f1254, plain, (spl66_110 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl66_110])])).
fof(f1246, plain, (spl66_109 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl66_109])])).
fof(f1238, plain, (spl66_108 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl66_108])])).
fof(f1230, plain, (spl66_107 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl66_107])])).
fof(f1222, plain, (spl66_106 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl66_106])])).
fof(f1214, plain, (spl66_105 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl66_105])])).
fof(f1206, plain, (spl66_104 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl66_104])])).
fof(f1198, plain, (spl66_103 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl66_103])])).
fof(f1190, plain, (spl66_102 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl66_102])])).
fof(f1182, plain, (spl66_101 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl66_101])])).
fof(f1174, plain, (spl66_100 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl66_100])])).
fof(f1166, plain, (spl66_99 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl66_99])])).
fof(f1158, plain, (spl66_98 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl66_98])])).
fof(f1150, plain, (spl66_97 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl66_97])])).
fof(f1142, plain, (spl66_96 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl66_96])])).
fof(f1134, plain, (spl66_95 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl66_95])])).
fof(f1126, plain, (spl66_94 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl66_94])])).
fof(f1118, plain, (spl66_93 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl66_93])])).
fof(f1110, plain, (spl66_92 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl66_92])])).
fof(f1102, plain, (spl66_91 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl66_91])])).
fof(f1094, plain, (spl66_90 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl66_90])])).
fof(f1086, plain, (spl66_89 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl66_89])])).
fof(f1078, plain, (spl66_88 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl66_88])])).
fof(f1070, plain, (spl66_87 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl66_87])])).
fof(f1062, plain, (spl66_86 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl66_86])])).
fof(f1054, plain, (spl66_85 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl66_85])])).
fof(f1046, plain, (spl66_84 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl66_84])])).
fof(f1038, plain, (spl66_83 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl66_83])])).
fof(f1030, plain, (spl66_82 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl66_82])])).
fof(f1022, plain, (spl66_81 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl66_81])])).
fof(f1014, plain, (spl66_80 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl66_80])])).
fof(f1006, plain, (spl66_79 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl66_79])])).
fof(f998, plain, (spl66_78 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl66_78])])).
fof(f990, plain, (spl66_77 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl66_77])])).
fof(f982, plain, (spl66_76 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl66_76])])).
fof(f974, plain, (spl66_75 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl66_75])])).
fof(f966, plain, (spl66_74 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl66_74])])).
fof(f958, plain, (spl66_73 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl66_73])])).
fof(f950, plain, (spl66_72 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl66_72])])).
fof(f942, plain, (spl66_71 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl66_71])])).
fof(f934, plain, (spl66_70 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl66_70])])).
fof(f926, plain, (spl66_69 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl66_69])])).
fof(f918, plain, (spl66_68 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl66_68])])).
fof(f910, plain, (spl66_67 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl66_67])])).
fof(f902, plain, (spl66_66 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl66_66])])).
fof(f894, plain, (spl66_65 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl66_65])])).
fof(f561, plain, ((e3 = op(e3, e3)) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f120])).
fof(f1421, plain, (spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | spl66_79 | spl66_78 | spl66_77 | spl66_76 | spl66_75 | spl66_74 | spl66_73 | spl66_72 | spl66_71 | spl66_70 | spl66_69 | spl66_68 | spl66_67 | spl66_66 | spl66_65 | ~ spl66_4), inference(avatar_split_clause, [], [f564, f602, f894, f902, f910, f918, f926, f934, f942, f950, f958, f966, f974, f982, f990, f998, f1006, f1014, f1022, f1030, f1038, f1046, f1054, f1062, f1070, f1078, f1086, f1094, f1102, f1110, f1118, f1126, f1134, f1142, f1150, f1158, f1166, f1174, f1182, f1190, f1198, f1206, f1214, f1222, f1230, f1238, f1246, f1254, f1262, f1270, f1278, f1286, f1294, f1302, f1310, f1318, f1326, f1334, f1342, f1350, f1358, f1366, f1374, f1382, f1390])).
fof(f564, plain, (~ (e3 = op(e3, e3)) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f120])).
fof(f1420, plain, (~ spl66_130 | ~ spl66_61), inference(avatar_split_clause, [], [f549, f845, f1414])).
fof(f549, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1419, plain, (~ spl66_130 | ~ spl66_46), inference(avatar_split_clause, [], [f550, f781, f1414])).
fof(f550, plain, (~ (e1 = op(e1, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1418, plain, (~ spl66_130 | ~ spl66_31), inference(avatar_split_clause, [], [f551, f717, f1414])).
fof(f551, plain, (~ (e2 = op(e2, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1417, plain, (~ spl66_130 | ~ spl66_16), inference(avatar_split_clause, [], [f552, f653, f1414])).
fof(f552, plain, (~ (e3 = op(e3, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1412, plain, (~ spl66_129 | ~ spl66_57), inference(avatar_split_clause, [], [f545, f828, f1406])).
fof(f545, plain, (~ (e0 = op(e0, e1)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1411, plain, (~ spl66_129 | ~ spl66_42), inference(avatar_split_clause, [], [f546, f764, f1406])).
fof(f546, plain, (~ (e1 = op(e1, e1)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1410, plain, (~ spl66_129 | ~ spl66_27), inference(avatar_split_clause, [], [f547, f700, f1406])).
fof(f547, plain, (~ (e2 = op(e2, e1)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1409, plain, (~ spl66_129 | ~ spl66_12), inference(avatar_split_clause, [], [f548, f636, f1406])).
fof(f548, plain, (~ (e3 = op(e3, e1)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1404, plain, (~ spl66_128 | ~ spl66_53), inference(avatar_split_clause, [], [f541, f811, f1398])).
fof(f541, plain, (~ (e0 = op(e0, e2)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1402, plain, (~ spl66_128 | ~ spl66_23), inference(avatar_split_clause, [], [f543, f683, f1398])).
fof(f543, plain, (~ (e2 = op(e2, e2)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1401, plain, (~ spl66_128 | ~ spl66_8), inference(avatar_split_clause, [], [f544, f619, f1398])).
fof(f544, plain, (~ (e3 = op(e3, e2)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1396, plain, (~ spl66_127 | spl66_61), inference(avatar_split_clause, [], [f537, f845, f1390])).
fof(f537, plain, ((e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1393, plain, (~ spl66_127 | ~ spl66_61), inference(avatar_split_clause, [], [f540, f845, f1390])).
fof(f540, plain, (~ (e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f1388, plain, (~ spl66_126 | spl66_61), inference(avatar_split_clause, [], [f533, f845, f1382])).
fof(f533, plain, ((e0 = op(e0, e0)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1386, plain, (~ spl66_126 | spl66_41), inference(avatar_split_clause, [], [f535, f760, f1382])).
fof(f535, plain, ((e0 = op(e1, e1)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f1385, plain, (~ spl66_126 | ~ spl66_46), inference(avatar_split_clause, [], [f536, f781, f1382])).
fof(f536, plain, (~ (e1 = op(e1, e0)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f1380, plain, (~ spl66_125 | spl66_61), inference(avatar_split_clause, [], [f529, f845, f1374])).
fof(f529, plain, ((e0 = op(e0, e0)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1378, plain, (~ spl66_125 | spl66_21), inference(avatar_split_clause, [], [f531, f675, f1374])).
fof(f531, plain, ((e0 = op(e2, e2)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f1377, plain, (~ spl66_125 | ~ spl66_31), inference(avatar_split_clause, [], [f532, f717, f1374])).
fof(f532, plain, (~ (e2 = op(e2, e0)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f1372, plain, (~ spl66_124 | spl66_61), inference(avatar_split_clause, [], [f525, f845, f1366])).
fof(f525, plain, ((e0 = op(e0, e0)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1370, plain, (~ spl66_124 | spl66_1), inference(avatar_split_clause, [], [f527, f590, f1366])).
fof(f527, plain, ((e0 = op(e3, e3)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1369, plain, (~ spl66_124 | ~ spl66_16), inference(avatar_split_clause, [], [f528, f653, f1366])).
fof(f528, plain, (~ (e3 = op(e3, e0)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1364, plain, (~ spl66_123 | spl66_57), inference(avatar_split_clause, [], [f521, f828, f1358])).
fof(f521, plain, ((e0 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1361, plain, (~ spl66_123 | ~ spl66_57), inference(avatar_split_clause, [], [f524, f828, f1358])).
fof(f524, plain, (~ (e0 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1356, plain, (~ spl66_122 | spl66_57), inference(avatar_split_clause, [], [f517, f828, f1350])).
fof(f517, plain, ((e0 = op(e0, e1)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1354, plain, (~ spl66_122 | spl66_42), inference(avatar_split_clause, [], [f519, f764, f1350])).
fof(f519, plain, ((e1 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f1353, plain, (~ spl66_122 | ~ spl66_42), inference(avatar_split_clause, [], [f520, f764, f1350])).
fof(f520, plain, (~ (e1 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f1348, plain, (~ spl66_121 | spl66_57), inference(avatar_split_clause, [], [f513, f828, f1342])).
fof(f513, plain, ((e0 = op(e0, e1)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1347, plain, (~ spl66_121 | spl66_45), inference(avatar_split_clause, [], [f514, f777, f1342])).
fof(f514, plain, ((e0 = op(e1, e0)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1346, plain, (~ spl66_121 | spl66_22), inference(avatar_split_clause, [], [f515, f679, f1342])).
fof(f515, plain, ((e1 = op(e2, e2)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1345, plain, (~ spl66_121 | ~ spl66_27), inference(avatar_split_clause, [], [f516, f700, f1342])).
fof(f516, plain, (~ (e2 = op(e2, e1)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1340, plain, (~ spl66_120 | spl66_57), inference(avatar_split_clause, [], [f509, f828, f1334])).
fof(f509, plain, ((e0 = op(e0, e1)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f1339, plain, (~ spl66_120 | spl66_45), inference(avatar_split_clause, [], [f510, f777, f1334])).
fof(f510, plain, ((e0 = op(e1, e0)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1338, plain, (~ spl66_120 | spl66_2), inference(avatar_split_clause, [], [f511, f594, f1334])).
fof(f511, plain, ((e1 = op(e3, e3)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1337, plain, (~ spl66_120 | ~ spl66_12), inference(avatar_split_clause, [], [f512, f636, f1334])).
fof(f512, plain, (~ (e3 = op(e3, e1)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1332, plain, (~ spl66_119 | spl66_53), inference(avatar_split_clause, [], [f505, f811, f1326])).
fof(f505, plain, ((e0 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f1329, plain, (~ spl66_119 | ~ spl66_53), inference(avatar_split_clause, [], [f508, f811, f1326])).
fof(f508, plain, (~ (e0 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1324, plain, (~ spl66_118 | spl66_53), inference(avatar_split_clause, [], [f501, f811, f1318])).
fof(f501, plain, ((e0 = op(e0, e2)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f1323, plain, (~ spl66_118 | spl66_29), inference(avatar_split_clause, [], [f502, f709, f1318])).
fof(f502, plain, ((e0 = op(e2, e0)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1322, plain, (~ spl66_118 | spl66_43), inference(avatar_split_clause, [], [f503, f768, f1318])).
fof(f503, plain, ((e2 = op(e1, e1)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1321, plain, (~ spl66_118 | ~ spl66_38), inference(avatar_split_clause, [], [f504, f747, f1318])).
fof(f504, plain, (~ (e1 = op(e1, e2)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1316, plain, (~ spl66_117 | spl66_53), inference(avatar_split_clause, [], [f497, f811, f1310])).
fof(f497, plain, ((e0 = op(e0, e2)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f1315, plain, (~ spl66_117 | spl66_29), inference(avatar_split_clause, [], [f498, f709, f1310])).
fof(f498, plain, ((e0 = op(e2, e0)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1314, plain, (~ spl66_117 | spl66_23), inference(avatar_split_clause, [], [f499, f683, f1310])).
fof(f499, plain, ((e2 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1313, plain, (~ spl66_117 | ~ spl66_23), inference(avatar_split_clause, [], [f500, f683, f1310])).
fof(f500, plain, (~ (e2 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1308, plain, (~ spl66_116 | spl66_53), inference(avatar_split_clause, [], [f493, f811, f1302])).
fof(f493, plain, ((e0 = op(e0, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f1307, plain, (~ spl66_116 | spl66_29), inference(avatar_split_clause, [], [f494, f709, f1302])).
fof(f494, plain, ((e0 = op(e2, e0)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1306, plain, (~ spl66_116 | spl66_3), inference(avatar_split_clause, [], [f495, f598, f1302])).
fof(f495, plain, ((e2 = op(e3, e3)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1305, plain, (~ spl66_116 | ~ spl66_8), inference(avatar_split_clause, [], [f496, f619, f1302])).
fof(f496, plain, (~ (e3 = op(e3, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1300, plain, (~ spl66_115 | spl66_49), inference(avatar_split_clause, [], [f489, f794, f1294])).
fof(f489, plain, ((e0 = op(e0, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f1298, plain, (~ spl66_115 | spl66_64), inference(avatar_split_clause, [], [f491, f857, f1294])).
fof(f491, plain, ((op(e0, e0) = e3) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1297, plain, (~ spl66_115 | ~ spl66_49), inference(avatar_split_clause, [], [f492, f794, f1294])).
fof(f492, plain, (~ (e0 = op(e0, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1292, plain, (~ spl66_114 | spl66_49), inference(avatar_split_clause, [], [f485, f794, f1286])).
fof(f485, plain, ((e0 = op(e0, e3)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f1291, plain, (~ spl66_114 | spl66_13), inference(avatar_split_clause, [], [f486, f641, f1286])).
fof(f486, plain, ((e0 = op(e3, e0)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1290, plain, (~ spl66_114 | spl66_44), inference(avatar_split_clause, [], [f487, f772, f1286])).
fof(f487, plain, ((e3 = op(e1, e1)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1289, plain, (~ spl66_114 | ~ spl66_34), inference(avatar_split_clause, [], [f488, f730, f1286])).
fof(f488, plain, (~ (e1 = op(e1, e3)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1284, plain, (~ spl66_113 | spl66_49), inference(avatar_split_clause, [], [f481, f794, f1278])).
fof(f481, plain, ((e0 = op(e0, e3)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f1283, plain, (~ spl66_113 | spl66_13), inference(avatar_split_clause, [], [f482, f641, f1278])).
fof(f482, plain, ((e0 = op(e3, e0)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1282, plain, (~ spl66_113 | spl66_24), inference(avatar_split_clause, [], [f483, f687, f1278])).
fof(f483, plain, ((e3 = op(e2, e2)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1281, plain, (~ spl66_113 | ~ spl66_19), inference(avatar_split_clause, [], [f484, f666, f1278])).
fof(f484, plain, (~ (e2 = op(e2, e3)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1276, plain, (~ spl66_112 | spl66_49), inference(avatar_split_clause, [], [f477, f794, f1270])).
fof(f477, plain, ((e0 = op(e0, e3)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f1274, plain, (~ spl66_112 | spl66_4), inference(avatar_split_clause, [], [f479, f602, f1270])).
fof(f479, plain, ((e3 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1273, plain, (~ spl66_112 | ~ spl66_4), inference(avatar_split_clause, [], [f480, f602, f1270])).
fof(f480, plain, (~ (e3 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1266, plain, (~ spl66_111 | spl66_61), inference(avatar_split_clause, [], [f475, f845, f1262])).
fof(f475, plain, ((e0 = op(e0, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f1265, plain, (~ spl66_111 | ~ spl66_61), inference(avatar_split_clause, [], [f476, f845, f1262])).
fof(f476, plain, (~ (e0 = op(e0, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1260, plain, (~ spl66_110 | spl66_46), inference(avatar_split_clause, [], [f469, f781, f1254])).
fof(f469, plain, ((e1 = op(e1, e0)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f1258, plain, (~ spl66_110 | spl66_41), inference(avatar_split_clause, [], [f471, f760, f1254])).
fof(f471, plain, ((e0 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1257, plain, (~ spl66_110 | ~ spl66_46), inference(avatar_split_clause, [], [f472, f781, f1254])).
fof(f472, plain, (~ (e1 = op(e1, e0)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1252, plain, (~ spl66_109 | spl66_46), inference(avatar_split_clause, [], [f465, f781, f1246])).
fof(f465, plain, ((e1 = op(e1, e0)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f1251, plain, (~ spl66_109 | spl66_58), inference(avatar_split_clause, [], [f466, f832, f1246])).
fof(f466, plain, ((e1 = op(e0, e1)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1250, plain, (~ spl66_109 | spl66_21), inference(avatar_split_clause, [], [f467, f675, f1246])).
fof(f467, plain, ((e0 = op(e2, e2)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1249, plain, (~ spl66_109 | ~ spl66_31), inference(avatar_split_clause, [], [f468, f717, f1246])).
fof(f468, plain, (~ (e2 = op(e2, e0)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1244, plain, (~ spl66_108 | spl66_46), inference(avatar_split_clause, [], [f461, f781, f1238])).
fof(f461, plain, ((e1 = op(e1, e0)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f1243, plain, (~ spl66_108 | spl66_58), inference(avatar_split_clause, [], [f462, f832, f1238])).
fof(f462, plain, ((e1 = op(e0, e1)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1242, plain, (~ spl66_108 | spl66_1), inference(avatar_split_clause, [], [f463, f590, f1238])).
fof(f463, plain, ((e0 = op(e3, e3)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1241, plain, (~ spl66_108 | ~ spl66_16), inference(avatar_split_clause, [], [f464, f653, f1238])).
fof(f464, plain, (~ (e3 = op(e3, e0)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1236, plain, (~ spl66_107 | spl66_42), inference(avatar_split_clause, [], [f457, f764, f1230])).
fof(f457, plain, ((e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f1234, plain, (~ spl66_107 | spl66_62), inference(avatar_split_clause, [], [f459, f849, f1230])).
fof(f459, plain, ((op(e0, e0) = e1) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f1233, plain, (~ spl66_107 | ~ spl66_57), inference(avatar_split_clause, [], [f460, f828, f1230])).
fof(f460, plain, (~ (e0 = op(e0, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f1228, plain, (~ spl66_106 | spl66_42), inference(avatar_split_clause, [], [f453, f764, f1222])).
fof(f453, plain, ((e1 = op(e1, e1)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f1225, plain, (~ spl66_106 | ~ spl66_42), inference(avatar_split_clause, [], [f456, f764, f1222])).
fof(f456, plain, (~ (e1 = op(e1, e1)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1220, plain, (~ spl66_105 | spl66_42), inference(avatar_split_clause, [], [f449, f764, f1214])).
fof(f449, plain, ((e1 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f1218, plain, (~ spl66_105 | spl66_22), inference(avatar_split_clause, [], [f451, f679, f1214])).
fof(f451, plain, ((e1 = op(e2, e2)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f1217, plain, (~ spl66_105 | ~ spl66_27), inference(avatar_split_clause, [], [f452, f700, f1214])).
fof(f452, plain, (~ (e2 = op(e2, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f1212, plain, (~ spl66_104 | spl66_42), inference(avatar_split_clause, [], [f445, f764, f1206])).
fof(f445, plain, ((e1 = op(e1, e1)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f1210, plain, (~ spl66_104 | spl66_2), inference(avatar_split_clause, [], [f447, f594, f1206])).
fof(f447, plain, ((e1 = op(e3, e3)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1209, plain, (~ spl66_104 | ~ spl66_12), inference(avatar_split_clause, [], [f448, f636, f1206])).
fof(f448, plain, (~ (e3 = op(e3, e1)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1204, plain, (~ spl66_103 | spl66_38), inference(avatar_split_clause, [], [f441, f747, f1198])).
fof(f441, plain, ((e1 = op(e1, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f1203, plain, (~ spl66_103 | spl66_26), inference(avatar_split_clause, [], [f442, f696, f1198])).
fof(f442, plain, ((e1 = op(e2, e1)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1202, plain, (~ spl66_103 | spl66_63), inference(avatar_split_clause, [], [f443, f853, f1198])).
fof(f443, plain, ((op(e0, e0) = e2) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1201, plain, (~ spl66_103 | ~ spl66_53), inference(avatar_split_clause, [], [f444, f811, f1198])).
fof(f444, plain, (~ (e0 = op(e0, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1196, plain, (~ spl66_102 | spl66_38), inference(avatar_split_clause, [], [f437, f747, f1190])).
fof(f437, plain, ((e1 = op(e1, e2)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f1194, plain, (~ spl66_102 | spl66_43), inference(avatar_split_clause, [], [f439, f768, f1190])).
fof(f439, plain, ((e2 = op(e1, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1193, plain, (~ spl66_102 | ~ spl66_38), inference(avatar_split_clause, [], [f440, f747, f1190])).
fof(f440, plain, (~ (e1 = op(e1, e2)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1187, plain, (~ spl66_101 | spl66_26), inference(avatar_split_clause, [], [f434, f696, f1182])).
fof(f434, plain, ((e1 = op(e2, e1)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f1186, plain, (~ spl66_101 | spl66_23), inference(avatar_split_clause, [], [f435, f683, f1182])).
fof(f435, plain, ((e2 = op(e2, e2)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f1185, plain, (~ spl66_101 | ~ spl66_23), inference(avatar_split_clause, [], [f436, f683, f1182])).
fof(f436, plain, (~ (e2 = op(e2, e2)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f1180, plain, (~ spl66_100 | spl66_38), inference(avatar_split_clause, [], [f429, f747, f1174])).
fof(f429, plain, ((e1 = op(e1, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP30), inference(nnf_transformation, [], [f84])).
fof(f1179, plain, (~ spl66_100 | spl66_26), inference(avatar_split_clause, [], [f430, f696, f1174])).
fof(f430, plain, ((e1 = op(e2, e1)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1178, plain, (~ spl66_100 | spl66_3), inference(avatar_split_clause, [], [f431, f598, f1174])).
fof(f431, plain, ((e2 = op(e3, e3)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1177, plain, (~ spl66_100 | ~ spl66_8), inference(avatar_split_clause, [], [f432, f619, f1174])).
fof(f432, plain, (~ (e3 = op(e3, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1172, plain, (~ spl66_99 | spl66_34), inference(avatar_split_clause, [], [f425, f730, f1166])).
fof(f425, plain, ((e1 = op(e1, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(nnf_transformation, [], [f85])).
fof(f1171, plain, (~ spl66_99 | spl66_10), inference(avatar_split_clause, [], [f426, f628, f1166])).
fof(f426, plain, ((e1 = op(e3, e1)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1170, plain, (~ spl66_99 | spl66_64), inference(avatar_split_clause, [], [f427, f857, f1166])).
fof(f427, plain, ((op(e0, e0) = e3) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1169, plain, (~ spl66_99 | ~ spl66_49), inference(avatar_split_clause, [], [f428, f794, f1166])).
fof(f428, plain, (~ (e0 = op(e0, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1164, plain, (~ spl66_98 | spl66_34), inference(avatar_split_clause, [], [f421, f730, f1158])).
fof(f421, plain, ((e1 = op(e1, e3)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP32), inference(nnf_transformation, [], [f86])).
fof(f1162, plain, (~ spl66_98 | spl66_44), inference(avatar_split_clause, [], [f423, f772, f1158])).
fof(f423, plain, ((e3 = op(e1, e1)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1161, plain, (~ spl66_98 | ~ spl66_34), inference(avatar_split_clause, [], [f424, f730, f1158])).
fof(f424, plain, (~ (e1 = op(e1, e3)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1156, plain, (~ spl66_97 | spl66_34), inference(avatar_split_clause, [], [f417, f730, f1150])).
fof(f417, plain, ((e1 = op(e1, e3)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP33), inference(nnf_transformation, [], [f87])).
fof(f1155, plain, (~ spl66_97 | spl66_10), inference(avatar_split_clause, [], [f418, f628, f1150])).
fof(f418, plain, ((e1 = op(e3, e1)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1154, plain, (~ spl66_97 | spl66_24), inference(avatar_split_clause, [], [f419, f687, f1150])).
fof(f419, plain, ((e3 = op(e2, e2)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1153, plain, (~ spl66_97 | ~ spl66_19), inference(avatar_split_clause, [], [f420, f666, f1150])).
fof(f420, plain, (~ (e2 = op(e2, e3)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1146, plain, (~ spl66_96 | spl66_4), inference(avatar_split_clause, [], [f415, f602, f1142])).
fof(f415, plain, ((e3 = op(e3, e3)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP34), inference(nnf_transformation, [], [f88])).
fof(f1145, plain, (~ spl66_96 | ~ spl66_4), inference(avatar_split_clause, [], [f416, f602, f1142])).
fof(f416, plain, (~ (e3 = op(e3, e3)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f1140, plain, (~ spl66_95 | spl66_31), inference(avatar_split_clause, [], [f409, f717, f1134])).
fof(f409, plain, ((e2 = op(e2, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(nnf_transformation, [], [f89])).
fof(f1138, plain, (~ spl66_95 | spl66_61), inference(avatar_split_clause, [], [f411, f845, f1134])).
fof(f411, plain, ((e0 = op(e0, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1137, plain, (~ spl66_95 | ~ spl66_61), inference(avatar_split_clause, [], [f412, f845, f1134])).
fof(f412, plain, (~ (e0 = op(e0, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1132, plain, (~ spl66_94 | spl66_31), inference(avatar_split_clause, [], [f405, f717, f1126])).
fof(f405, plain, ((e2 = op(e2, e0)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP36), inference(nnf_transformation, [], [f90])).
fof(f1131, plain, (~ spl66_94 | spl66_55), inference(avatar_split_clause, [], [f406, f819, f1126])).
fof(f406, plain, ((e2 = op(e0, e2)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1130, plain, (~ spl66_94 | spl66_41), inference(avatar_split_clause, [], [f407, f760, f1126])).
fof(f407, plain, ((e0 = op(e1, e1)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1129, plain, (~ spl66_94 | ~ spl66_46), inference(avatar_split_clause, [], [f408, f781, f1126])).
fof(f408, plain, (~ (e1 = op(e1, e0)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1124, plain, (~ spl66_93 | spl66_31), inference(avatar_split_clause, [], [f401, f717, f1118])).
fof(f401, plain, ((e2 = op(e2, e0)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP37), inference(nnf_transformation, [], [f91])).
fof(f1122, plain, (~ spl66_93 | spl66_21), inference(avatar_split_clause, [], [f403, f675, f1118])).
fof(f403, plain, ((e0 = op(e2, e2)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1121, plain, (~ spl66_93 | ~ spl66_31), inference(avatar_split_clause, [], [f404, f717, f1118])).
fof(f404, plain, (~ (e2 = op(e2, e0)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1116, plain, (~ spl66_92 | spl66_31), inference(avatar_split_clause, [], [f397, f717, f1110])).
fof(f397, plain, ((e2 = op(e2, e0)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP38), inference(nnf_transformation, [], [f92])).
fof(f1115, plain, (~ spl66_92 | spl66_55), inference(avatar_split_clause, [], [f398, f819, f1110])).
fof(f398, plain, ((e2 = op(e0, e2)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1114, plain, (~ spl66_92 | spl66_1), inference(avatar_split_clause, [], [f399, f590, f1110])).
fof(f399, plain, ((e0 = op(e3, e3)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1113, plain, (~ spl66_92 | ~ spl66_16), inference(avatar_split_clause, [], [f400, f653, f1110])).
fof(f400, plain, (~ (e3 = op(e3, e0)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1108, plain, (~ spl66_91 | spl66_27), inference(avatar_split_clause, [], [f393, f700, f1102])).
fof(f393, plain, ((e2 = op(e2, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(nnf_transformation, [], [f93])).
fof(f1107, plain, (~ spl66_91 | spl66_39), inference(avatar_split_clause, [], [f394, f751, f1102])).
fof(f394, plain, ((e2 = op(e1, e2)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1106, plain, (~ spl66_91 | spl66_62), inference(avatar_split_clause, [], [f395, f849, f1102])).
fof(f395, plain, ((op(e0, e0) = e1) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1105, plain, (~ spl66_91 | ~ spl66_57), inference(avatar_split_clause, [], [f396, f828, f1102])).
fof(f396, plain, (~ (e0 = op(e0, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1098, plain, (~ spl66_90 | spl66_42), inference(avatar_split_clause, [], [f391, f764, f1094])).
fof(f391, plain, ((e1 = op(e1, e1)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP40), inference(nnf_transformation, [], [f94])).
fof(f1097, plain, (~ spl66_90 | ~ spl66_42), inference(avatar_split_clause, [], [f392, f764, f1094])).
fof(f392, plain, (~ (e1 = op(e1, e1)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f1092, plain, (~ spl66_89 | spl66_27), inference(avatar_split_clause, [], [f385, f700, f1086])).
fof(f385, plain, ((e2 = op(e2, e1)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP41), inference(nnf_transformation, [], [f95])).
fof(f1089, plain, (~ spl66_89 | ~ spl66_27), inference(avatar_split_clause, [], [f388, f700, f1086])).
fof(f388, plain, (~ (e2 = op(e2, e1)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1084, plain, (~ spl66_88 | spl66_27), inference(avatar_split_clause, [], [f381, f700, f1078])).
fof(f381, plain, ((e2 = op(e2, e1)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP42), inference(nnf_transformation, [], [f96])).
fof(f1083, plain, (~ spl66_88 | spl66_39), inference(avatar_split_clause, [], [f382, f751, f1078])).
fof(f382, plain, ((e2 = op(e1, e2)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1082, plain, (~ spl66_88 | spl66_2), inference(avatar_split_clause, [], [f383, f594, f1078])).
fof(f383, plain, ((e1 = op(e3, e3)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1081, plain, (~ spl66_88 | ~ spl66_12), inference(avatar_split_clause, [], [f384, f636, f1078])).
fof(f384, plain, (~ (e3 = op(e3, e1)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1076, plain, (~ spl66_87 | spl66_23), inference(avatar_split_clause, [], [f377, f683, f1070])).
fof(f377, plain, ((e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(nnf_transformation, [], [f97])).
fof(f1074, plain, (~ spl66_87 | spl66_63), inference(avatar_split_clause, [], [f379, f853, f1070])).
fof(f379, plain, ((op(e0, e0) = e2) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1073, plain, (~ spl66_87 | ~ spl66_53), inference(avatar_split_clause, [], [f380, f811, f1070])).
fof(f380, plain, (~ (e0 = op(e0, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1068, plain, (~ spl66_86 | spl66_23), inference(avatar_split_clause, [], [f373, f683, f1062])).
fof(f373, plain, ((e2 = op(e2, e2)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP44), inference(nnf_transformation, [], [f98])).
fof(f1066, plain, (~ spl66_86 | spl66_43), inference(avatar_split_clause, [], [f375, f768, f1062])).
fof(f375, plain, ((e2 = op(e1, e1)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1065, plain, (~ spl66_86 | ~ spl66_38), inference(avatar_split_clause, [], [f376, f747, f1062])).
fof(f376, plain, (~ (e1 = op(e1, e2)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1060, plain, (~ spl66_85 | spl66_23), inference(avatar_split_clause, [], [f369, f683, f1054])).
fof(f369, plain, ((e2 = op(e2, e2)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP45), inference(nnf_transformation, [], [f99])).
fof(f1057, plain, (~ spl66_85 | ~ spl66_23), inference(avatar_split_clause, [], [f372, f683, f1054])).
fof(f372, plain, (~ (e2 = op(e2, e2)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1052, plain, (~ spl66_84 | spl66_23), inference(avatar_split_clause, [], [f365, f683, f1046])).
fof(f365, plain, ((e2 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP46), inference(nnf_transformation, [], [f100])).
fof(f1050, plain, (~ spl66_84 | spl66_3), inference(avatar_split_clause, [], [f367, f598, f1046])).
fof(f367, plain, ((e2 = op(e3, e3)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f1049, plain, (~ spl66_84 | ~ spl66_8), inference(avatar_split_clause, [], [f368, f619, f1046])).
fof(f368, plain, (~ (e3 = op(e3, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f1044, plain, (~ spl66_83 | spl66_19), inference(avatar_split_clause, [], [f361, f666, f1038])).
fof(f361, plain, ((e2 = op(e2, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(nnf_transformation, [], [f101])).
fof(f1043, plain, (~ spl66_83 | spl66_7), inference(avatar_split_clause, [], [f362, f615, f1038])).
fof(f362, plain, ((e2 = op(e3, e2)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1042, plain, (~ spl66_83 | spl66_64), inference(avatar_split_clause, [], [f363, f857, f1038])).
fof(f363, plain, ((op(e0, e0) = e3) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1041, plain, (~ spl66_83 | ~ spl66_49), inference(avatar_split_clause, [], [f364, f794, f1038])).
fof(f364, plain, (~ (e0 = op(e0, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1036, plain, (~ spl66_82 | spl66_19), inference(avatar_split_clause, [], [f357, f666, f1030])).
fof(f357, plain, ((e2 = op(e2, e3)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP48), inference(nnf_transformation, [], [f102])).
fof(f1035, plain, (~ spl66_82 | spl66_7), inference(avatar_split_clause, [], [f358, f615, f1030])).
fof(f358, plain, ((e2 = op(e3, e2)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1034, plain, (~ spl66_82 | spl66_44), inference(avatar_split_clause, [], [f359, f772, f1030])).
fof(f359, plain, ((e3 = op(e1, e1)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1033, plain, (~ spl66_82 | ~ spl66_34), inference(avatar_split_clause, [], [f360, f730, f1030])).
fof(f360, plain, (~ (e1 = op(e1, e3)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1028, plain, (~ spl66_81 | spl66_19), inference(avatar_split_clause, [], [f353, f666, f1022])).
fof(f353, plain, ((e2 = op(e2, e3)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP49), inference(nnf_transformation, [], [f103])).
fof(f1025, plain, (~ spl66_81 | ~ spl66_19), inference(avatar_split_clause, [], [f356, f666, f1022])).
fof(f356, plain, (~ (e2 = op(e2, e3)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1020, plain, (~ spl66_80 | spl66_19), inference(avatar_split_clause, [], [f349, f666, f1014])).
fof(f349, plain, ((e2 = op(e2, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP50), inference(nnf_transformation, [], [f104])).
fof(f1019, plain, (~ spl66_80 | spl66_7), inference(avatar_split_clause, [], [f350, f615, f1014])).
fof(f350, plain, ((e2 = op(e3, e2)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1018, plain, (~ spl66_80 | spl66_4), inference(avatar_split_clause, [], [f351, f602, f1014])).
fof(f351, plain, ((e3 = op(e3, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1017, plain, (~ spl66_80 | ~ spl66_4), inference(avatar_split_clause, [], [f352, f602, f1014])).
fof(f352, plain, (~ (e3 = op(e3, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1012, plain, (~ spl66_79 | spl66_16), inference(avatar_split_clause, [], [f345, f653, f1006])).
fof(f345, plain, ((e3 = op(e3, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(nnf_transformation, [], [f105])).
fof(f1010, plain, (~ spl66_79 | spl66_61), inference(avatar_split_clause, [], [f347, f845, f1006])).
fof(f347, plain, ((e0 = op(e0, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f1009, plain, (~ spl66_79 | ~ spl66_61), inference(avatar_split_clause, [], [f348, f845, f1006])).
fof(f348, plain, (~ (e0 = op(e0, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f1004, plain, (~ spl66_78 | spl66_16), inference(avatar_split_clause, [], [f341, f653, f998])).
fof(f341, plain, ((e3 = op(e3, e0)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP52), inference(nnf_transformation, [], [f106])).
fof(f1003, plain, (~ spl66_78 | spl66_52), inference(avatar_split_clause, [], [f342, f806, f998])).
fof(f342, plain, ((e3 = op(e0, e3)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1002, plain, (~ spl66_78 | spl66_41), inference(avatar_split_clause, [], [f343, f760, f998])).
fof(f343, plain, ((e0 = op(e1, e1)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1001, plain, (~ spl66_78 | ~ spl66_46), inference(avatar_split_clause, [], [f344, f781, f998])).
fof(f344, plain, (~ (e1 = op(e1, e0)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f996, plain, (~ spl66_77 | spl66_16), inference(avatar_split_clause, [], [f337, f653, f990])).
fof(f337, plain, ((e3 = op(e3, e0)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP53), inference(nnf_transformation, [], [f107])).
fof(f995, plain, (~ spl66_77 | spl66_52), inference(avatar_split_clause, [], [f338, f806, f990])).
fof(f338, plain, ((e3 = op(e0, e3)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f994, plain, (~ spl66_77 | spl66_21), inference(avatar_split_clause, [], [f339, f675, f990])).
fof(f339, plain, ((e0 = op(e2, e2)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f993, plain, (~ spl66_77 | ~ spl66_31), inference(avatar_split_clause, [], [f340, f717, f990])).
fof(f340, plain, (~ (e2 = op(e2, e0)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f988, plain, (~ spl66_76 | spl66_16), inference(avatar_split_clause, [], [f333, f653, f982])).
fof(f333, plain, ((e3 = op(e3, e0)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP54), inference(nnf_transformation, [], [f108])).
fof(f987, plain, (~ spl66_76 | spl66_52), inference(avatar_split_clause, [], [f334, f806, f982])).
fof(f334, plain, ((e3 = op(e0, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f986, plain, (~ spl66_76 | spl66_1), inference(avatar_split_clause, [], [f335, f590, f982])).
fof(f335, plain, ((e0 = op(e3, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f985, plain, (~ spl66_76 | ~ spl66_16), inference(avatar_split_clause, [], [f336, f653, f982])).
fof(f336, plain, (~ (e3 = op(e3, e0)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f980, plain, (~ spl66_75 | spl66_12), inference(avatar_split_clause, [], [f329, f636, f974])).
fof(f329, plain, ((e3 = op(e3, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(nnf_transformation, [], [f109])).
fof(f979, plain, (~ spl66_75 | spl66_36), inference(avatar_split_clause, [], [f330, f738, f974])).
fof(f330, plain, ((e3 = op(e1, e3)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f978, plain, (~ spl66_75 | spl66_62), inference(avatar_split_clause, [], [f331, f849, f974])).
fof(f331, plain, ((op(e0, e0) = e1) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f977, plain, (~ spl66_75 | ~ spl66_57), inference(avatar_split_clause, [], [f332, f828, f974])).
fof(f332, plain, (~ (e0 = op(e0, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f972, plain, (~ spl66_74 | spl66_12), inference(avatar_split_clause, [], [f325, f636, f966])).
fof(f325, plain, ((e3 = op(e3, e1)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP56), inference(nnf_transformation, [], [f110])).
fof(f970, plain, (~ spl66_74 | spl66_42), inference(avatar_split_clause, [], [f327, f764, f966])).
fof(f327, plain, ((e1 = op(e1, e1)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f969, plain, (~ spl66_74 | ~ spl66_42), inference(avatar_split_clause, [], [f328, f764, f966])).
fof(f328, plain, (~ (e1 = op(e1, e1)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f964, plain, (~ spl66_73 | spl66_12), inference(avatar_split_clause, [], [f321, f636, f958])).
fof(f321, plain, ((e3 = op(e3, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP57), inference(nnf_transformation, [], [f111])).
fof(f963, plain, (~ spl66_73 | spl66_36), inference(avatar_split_clause, [], [f322, f738, f958])).
fof(f322, plain, ((e3 = op(e1, e3)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f962, plain, (~ spl66_73 | spl66_22), inference(avatar_split_clause, [], [f323, f679, f958])).
fof(f323, plain, ((e1 = op(e2, e2)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f961, plain, (~ spl66_73 | ~ spl66_27), inference(avatar_split_clause, [], [f324, f700, f958])).
fof(f324, plain, (~ (e2 = op(e2, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f956, plain, (~ spl66_72 | spl66_12), inference(avatar_split_clause, [], [f317, f636, f950])).
fof(f317, plain, ((e3 = op(e3, e1)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP58), inference(nnf_transformation, [], [f112])).
fof(f953, plain, (~ spl66_72 | ~ spl66_12), inference(avatar_split_clause, [], [f320, f636, f950])).
fof(f320, plain, (~ (e3 = op(e3, e1)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f948, plain, (~ spl66_71 | spl66_8), inference(avatar_split_clause, [], [f313, f619, f942])).
fof(f313, plain, ((e3 = op(e3, e2)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(nnf_transformation, [], [f113])).
fof(f947, plain, (~ spl66_71 | spl66_20), inference(avatar_split_clause, [], [f314, f670, f942])).
fof(f314, plain, ((e3 = op(e2, e3)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f946, plain, (~ spl66_71 | spl66_63), inference(avatar_split_clause, [], [f315, f853, f942])).
fof(f315, plain, ((op(e0, e0) = e2) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f945, plain, (~ spl66_71 | ~ spl66_53), inference(avatar_split_clause, [], [f316, f811, f942])).
fof(f316, plain, (~ (e0 = op(e0, e2)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f940, plain, (~ spl66_70 | spl66_8), inference(avatar_split_clause, [], [f309, f619, f934])).
fof(f309, plain, ((e3 = op(e3, e2)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP60), inference(nnf_transformation, [], [f114])).
fof(f939, plain, (~ spl66_70 | spl66_20), inference(avatar_split_clause, [], [f310, f670, f934])).
fof(f310, plain, ((e3 = op(e2, e3)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f938, plain, (~ spl66_70 | spl66_43), inference(avatar_split_clause, [], [f311, f768, f934])).
fof(f311, plain, ((e2 = op(e1, e1)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f937, plain, (~ spl66_70 | ~ spl66_38), inference(avatar_split_clause, [], [f312, f747, f934])).
fof(f312, plain, (~ (e1 = op(e1, e2)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f932, plain, (~ spl66_69 | spl66_8), inference(avatar_split_clause, [], [f305, f619, f926])).
fof(f305, plain, ((e3 = op(e3, e2)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP61), inference(nnf_transformation, [], [f115])).
fof(f931, plain, (~ spl66_69 | spl66_20), inference(avatar_split_clause, [], [f306, f670, f926])).
fof(f306, plain, ((e3 = op(e2, e3)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f930, plain, (~ spl66_69 | spl66_23), inference(avatar_split_clause, [], [f307, f683, f926])).
fof(f307, plain, ((e2 = op(e2, e2)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f929, plain, (~ spl66_69 | ~ spl66_23), inference(avatar_split_clause, [], [f308, f683, f926])).
fof(f308, plain, (~ (e2 = op(e2, e2)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f924, plain, (~ spl66_68 | spl66_8), inference(avatar_split_clause, [], [f301, f619, f918])).
fof(f301, plain, ((e3 = op(e3, e2)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP62), inference(nnf_transformation, [], [f116])).
fof(f922, plain, (~ spl66_68 | spl66_3), inference(avatar_split_clause, [], [f303, f598, f918])).
fof(f303, plain, ((e2 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f921, plain, (~ spl66_68 | ~ spl66_8), inference(avatar_split_clause, [], [f304, f619, f918])).
fof(f304, plain, (~ (e3 = op(e3, e2)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f916, plain, (~ spl66_67 | spl66_4), inference(avatar_split_clause, [], [f297, f602, f910])).
fof(f297, plain, ((e3 = op(e3, e3)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP63), inference(nnf_transformation, [], [f117])).
fof(f914, plain, (~ spl66_67 | spl66_64), inference(avatar_split_clause, [], [f299, f857, f910])).
fof(f299, plain, ((op(e0, e0) = e3) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f913, plain, (~ spl66_67 | ~ spl66_49), inference(avatar_split_clause, [], [f300, f794, f910])).
fof(f300, plain, (~ (e0 = op(e0, e3)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f908, plain, (~ spl66_66 | spl66_4), inference(avatar_split_clause, [], [f293, f602, f902])).
fof(f293, plain, ((e3 = op(e3, e3)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP64), inference(nnf_transformation, [], [f118])).
fof(f906, plain, (~ spl66_66 | spl66_44), inference(avatar_split_clause, [], [f295, f772, f902])).
fof(f295, plain, ((e3 = op(e1, e1)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f905, plain, (~ spl66_66 | ~ spl66_34), inference(avatar_split_clause, [], [f296, f730, f902])).
fof(f296, plain, (~ (e1 = op(e1, e3)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f900, plain, (~ spl66_65 | spl66_4), inference(avatar_split_clause, [], [f289, f602, f894])).
fof(f289, plain, ((e3 = op(e3, e3)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP65), inference(nnf_transformation, [], [f119])).
fof(f898, plain, (~ spl66_65 | spl66_24), inference(avatar_split_clause, [], [f291, f687, f894])).
fof(f291, plain, ((e3 = op(e2, e2)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f897, plain, (~ spl66_65 | ~ spl66_19), inference(avatar_split_clause, [], [f292, f666, f894])).
fof(f292, plain, (~ (e2 = op(e2, e3)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f892, plain, (spl66_61 | spl66_57 | spl66_53 | spl66_49), inference(avatar_split_clause, [], [f203, f794, f811, f828, f845])).
fof(f203, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax2)).
fof(f891, plain, (spl66_61 | spl66_45 | spl66_29 | spl66_13), inference(avatar_split_clause, [], [f204, f641, f709, f777, f845])).
fof(f204, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f890, plain, (spl66_62 | spl66_58 | spl66_54 | spl66_50), inference(avatar_split_clause, [], [f205, f798, f815, f832, f849])).
fof(f205, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f889, plain, (spl66_62 | spl66_46 | spl66_30 | spl66_14), inference(avatar_split_clause, [], [f206, f645, f713, f781, f849])).
fof(f206, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f888, plain, (spl66_63 | spl66_59 | spl66_55 | spl66_51), inference(avatar_split_clause, [], [f207, f802, f819, f836, f853])).
fof(f207, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f887, plain, (spl66_63 | spl66_47 | spl66_31 | spl66_15), inference(avatar_split_clause, [], [f208, f649, f717, f785, f853])).
fof(f208, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f886, plain, (spl66_64 | spl66_60 | spl66_56 | spl66_52), inference(avatar_split_clause, [], [f209, f806, f823, f840, f857])).
fof(f209, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f885, plain, (spl66_64 | spl66_48 | spl66_32 | spl66_16), inference(avatar_split_clause, [], [f210, f653, f721, f789, f857])).
fof(f210, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f884, plain, (spl66_45 | spl66_41 | spl66_37 | spl66_33), inference(avatar_split_clause, [], [f211, f726, f743, f760, f777])).
fof(f211, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f883, plain, (spl66_57 | spl66_41 | spl66_25 | spl66_9), inference(avatar_split_clause, [], [f212, f624, f692, f760, f828])).
fof(f212, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f882, plain, (spl66_46 | spl66_42 | spl66_38 | spl66_34), inference(avatar_split_clause, [], [f213, f730, f747, f764, f781])).
fof(f213, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f880, plain, (spl66_47 | spl66_43 | spl66_39 | spl66_35), inference(avatar_split_clause, [], [f215, f734, f751, f768, f785])).
fof(f215, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f879, plain, (spl66_59 | spl66_43 | spl66_27 | spl66_11), inference(avatar_split_clause, [], [f216, f632, f700, f768, f836])).
fof(f216, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f878, plain, (spl66_48 | spl66_44 | spl66_40 | spl66_36), inference(avatar_split_clause, [], [f217, f738, f755, f772, f789])).
fof(f217, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f877, plain, (spl66_60 | spl66_44 | spl66_28 | spl66_12), inference(avatar_split_clause, [], [f218, f636, f704, f772, f840])).
fof(f218, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f876, plain, (spl66_29 | spl66_25 | spl66_21 | spl66_17), inference(avatar_split_clause, [], [f219, f658, f675, f692, f709])).
fof(f219, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f875, plain, (spl66_53 | spl66_37 | spl66_21 | spl66_5), inference(avatar_split_clause, [], [f220, f607, f675, f743, f811])).
fof(f220, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f874, plain, (spl66_30 | spl66_26 | spl66_22 | spl66_18), inference(avatar_split_clause, [], [f221, f662, f679, f696, f713])).
fof(f221, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f873, plain, (spl66_54 | spl66_38 | spl66_22 | spl66_6), inference(avatar_split_clause, [], [f222, f611, f679, f747, f815])).
fof(f222, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f872, plain, (spl66_31 | spl66_27 | spl66_23 | spl66_19), inference(avatar_split_clause, [], [f223, f666, f683, f700, f717])).
fof(f223, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f871, plain, (spl66_55 | spl66_39 | spl66_23 | spl66_7), inference(avatar_split_clause, [], [f224, f615, f683, f751, f819])).
fof(f224, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f870, plain, (spl66_32 | spl66_28 | spl66_24 | spl66_20), inference(avatar_split_clause, [], [f225, f670, f687, f704, f721])).
fof(f225, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f869, plain, (spl66_56 | spl66_40 | spl66_24 | spl66_8), inference(avatar_split_clause, [], [f226, f619, f687, f755, f823])).
fof(f226, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f868, plain, (spl66_13 | spl66_9 | spl66_5 | spl66_1), inference(avatar_split_clause, [], [f227, f590, f607, f624, f641])).
fof(f227, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f867, plain, (spl66_49 | spl66_33 | spl66_17 | spl66_1), inference(avatar_split_clause, [], [f228, f590, f658, f726, f794])).
fof(f228, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f866, plain, (spl66_14 | spl66_10 | spl66_6 | spl66_2), inference(avatar_split_clause, [], [f229, f594, f611, f628, f645])).
fof(f229, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f865, plain, (spl66_50 | spl66_34 | spl66_18 | spl66_2), inference(avatar_split_clause, [], [f230, f594, f662, f730, f798])).
fof(f230, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f864, plain, (spl66_15 | spl66_11 | spl66_7 | spl66_3), inference(avatar_split_clause, [], [f231, f598, f615, f632, f649])).
fof(f231, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f863, plain, (spl66_51 | spl66_35 | spl66_19 | spl66_3), inference(avatar_split_clause, [], [f232, f598, f666, f734, f802])).
fof(f232, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f862, plain, (spl66_16 | spl66_12 | spl66_8 | spl66_4), inference(avatar_split_clause, [], [f233, f602, f619, f636, f653])).
fof(f233, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f861, plain, (spl66_52 | spl66_36 | spl66_20 | spl66_4), inference(avatar_split_clause, [], [f234, f602, f670, f738, f806])).
fof(f234, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f860, plain, (spl66_61 | spl66_62 | spl66_63 | spl66_64), inference(avatar_split_clause, [], [f187, f857, f853, f849, f845])).
fof(f187, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG144+1.p', ax1)).
fof(f843, plain, (spl66_57 | spl66_58 | spl66_59 | spl66_60), inference(avatar_split_clause, [], [f188, f840, f836, f832, f828])).
fof(f188, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f826, plain, (spl66_53 | spl66_54 | spl66_55 | spl66_56), inference(avatar_split_clause, [], [f189, f823, f819, f815, f811])).
fof(f189, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f809, plain, (spl66_49 | spl66_50 | spl66_51 | spl66_52), inference(avatar_split_clause, [], [f190, f806, f802, f798, f794])).
fof(f190, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f792, plain, (spl66_45 | spl66_46 | spl66_47 | spl66_48), inference(avatar_split_clause, [], [f191, f789, f785, f781, f777])).
fof(f191, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f775, plain, (spl66_41 | spl66_42 | spl66_43 | spl66_44), inference(avatar_split_clause, [], [f192, f772, f768, f764, f760])).
fof(f192, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f758, plain, (spl66_37 | spl66_38 | spl66_39 | spl66_40), inference(avatar_split_clause, [], [f193, f755, f751, f747, f743])).
fof(f193, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f741, plain, (spl66_33 | spl66_34 | spl66_35 | spl66_36), inference(avatar_split_clause, [], [f194, f738, f734, f730, f726])).
fof(f194, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f724, plain, (spl66_29 | spl66_30 | spl66_31 | spl66_32), inference(avatar_split_clause, [], [f195, f721, f717, f713, f709])).
fof(f195, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f707, plain, (spl66_25 | spl66_26 | spl66_27 | spl66_28), inference(avatar_split_clause, [], [f196, f704, f700, f696, f692])).
fof(f196, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f690, plain, (spl66_21 | spl66_22 | spl66_23 | spl66_24), inference(avatar_split_clause, [], [f197, f687, f683, f679, f675])).
fof(f197, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f673, plain, (spl66_17 | spl66_18 | spl66_19 | spl66_20), inference(avatar_split_clause, [], [f198, f670, f666, f662, f658])).
fof(f198, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f656, plain, (spl66_13 | spl66_14 | spl66_15 | spl66_16), inference(avatar_split_clause, [], [f199, f653, f649, f645, f641])).
fof(f199, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f639, plain, (spl66_9 | spl66_10 | spl66_11 | spl66_12), inference(avatar_split_clause, [], [f200, f636, f632, f628, f624])).
fof(f200, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f622, plain, (spl66_5 | spl66_6 | spl66_7 | spl66_8), inference(avatar_split_clause, [], [f201, f619, f615, f611, f607])).
fof(f201, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f605, plain, (spl66_1 | spl66_2 | spl66_3 | spl66_4), inference(avatar_split_clause, [], [f202, f602, f598, f594, f590])).
fof(f202, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).