fof(f1944, plain, $false, inference(avatar_sat_refutation, [], [f858, f1366, f1368, f1667, f1858])).
fof(f1858, plain, (spl4_8 | ~ spl4_40 | ~ spl4_59), inference(avatar_split_clause, [], [f1854, f527, f447, f313])).
fof(f313, plain, (spl4_8 <=> (e2 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_8])])).
fof(f447, plain, (spl4_40 <=> (e4 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl4_40])])).
fof(f527, plain, (spl4_59 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_59])])).
fof(f1854, plain, ((e2 = op(e4, e3)) | (~ spl4_40 | ~ spl4_59)), inference(backward_demodulation, [], [f1699, f529])).
fof(f529, plain, ((e3 = op(e2, e3)) | ~ spl4_59), inference(avatar_component_clause, [], [f527])).
fof(f1699, plain, ((e2 = op(e4, op(e2, e3))) | ~ spl4_40), inference(forward_demodulation, [], [f107, f449])).
fof(f449, plain, ((e4 = op(e3, e2)) | ~ spl4_40), inference(avatar_component_clause, [], [f447])).
fof(f107, plain, (e2 = op(op(e3, e2), op(e2, e3))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, ((e4 = op(op(e4, e4), op(e4, e4))) & (e4 = op(op(e3, e4), op(e4, e3))) & (e4 = op(op(e2, e4), op(e4, e2))) & (e4 = op(op(e1, e4), op(e4, e1))) & (e4 = op(op(e0, e4), op(e4, e0))) & (e3 = op(op(e4, e3), op(e3, e4))) & (e3 = op(op(e3, e3), op(e3, e3))) & (e3 = op(op(e2, e3), op(e3, e2))) & (e3 = op(op(e1, e3), op(e3, e1))) & (e3 = op(op(e0, e3), op(e3, e0))) & (e2 = op(op(e4, e2), op(e2, e4))) & (e2 = op(op(e3, e2), op(e2, e3))) & (e2 = op(op(e2, e2), op(e2, e2))) & (e2 = op(op(e1, e2), op(e2, e1))) & (e2 = op(op(e0, e2), op(e2, e0))) & (e1 = op(op(e4, e1), op(e1, e4))) & (e1 = op(op(e3, e1), op(e1, e3))) & (e1 = op(op(e2, e1), op(e1, e2))) & (e1 = op(op(e1, e1), op(e1, e1))) & (e1 = op(op(e0, e1), op(e1, e0))) & (e0 = op(op(e4, e0), op(e0, e4))) & (e0 = op(op(e3, e0), op(e0, e3))) & (e0 = op(op(e2, e0), op(e0, e2))) & (e0 = op(op(e1, e0), op(e0, e1))) & (e0 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG174+1.p', ax3)).
fof(f1667, plain, (~ spl4_8 | ~ spl4_33), inference(avatar_split_clause, [], [f1666, f418, f313])).
fof(f418, plain, (spl4_33 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl4_33])])).
fof(f1666, plain, (~ (e2 = op(e4, e3)) | ~ spl4_33), inference(forward_demodulation, [], [f158, f420])).
fof(f420, plain, ((e2 = op(e3, e3)) | ~ spl4_33), inference(avatar_component_clause, [], [f418])).
fof(f158, plain, ~ (op(e3, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (op(e4, e3) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e4)) & ~ (op(e4, e1) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e4)) & ~ (op(e4, e0) = op(e4, e3)) & ~ (op(e4, e0) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e1)) & ~ (op(e3, e3) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e4)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e4)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e3) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e4)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e4)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e3) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e4)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e4)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e3) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e4)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e4)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e3, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e4, e4)) & ~ (op(e1, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e4, e4)) & ~ (op(e0, e4) = op(e3, e4)) & ~ (op(e0, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e1, e4)) & ~ (op(e3, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e4, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e4, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e3, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e4, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e4, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e3, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e4, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e4, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e3, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e4, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e4, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG174+1.p', ax4)).
fof(f1368, plain, (spl4_40 | ~ spl4_33), inference(avatar_split_clause, [], [f1367, f418, f447])).
fof(f1367, plain, ((e4 = op(e3, e2)) | ~ spl4_33), inference(forward_demodulation, [], [f232, f420])).
fof(f232, plain, (e4 = op(e3, op(e3, e3))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ((e4 = op(e3, op(e3, e3))) & (e2 = op(e3, e3)) & (e1 = op(e3, op(e3, op(e3, e3)))) & (e0 = op(op(e3, op(e3, e3)), op(e3, op(e3, e3))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG174+1.p', ax6)).
fof(f1366, plain, (spl4_59 | ~ spl4_33), inference(avatar_split_clause, [], [f1365, f418, f527])).
fof(f1365, plain, ((e3 = op(e2, e3)) | ~ spl4_33), inference(forward_demodulation, [], [f276, f420])).
fof(f276, plain, (e3 = op(op(e3, e3), e3)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((e4 = op(op(e4, e4), e4)) & (e4 = op(op(e4, e3), e3)) & (e4 = op(op(e4, e2), e2)) & (e4 = op(op(e4, e1), e1)) & (e4 = op(op(e4, e0), e0)) & (e3 = op(op(e3, e4), e4)) & (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e2), e2)) & (e3 = op(op(e3, e1), e1)) & (e3 = op(op(e3, e0), e0)) & (e2 = op(op(e2, e4), e4)) & (e2 = op(op(e2, e3), e3)) & (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e1), e1)) & (e2 = op(op(e2, e0), e0)) & (e1 = op(op(e1, e4), e4)) & (e1 = op(op(e1, e3), e3)) & (e1 = op(op(e1, e2), e2)) & (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e0, e4), e4)) & (e0 = op(op(e0, e3), e3)) & (e0 = op(op(e0, e2), e2)) & (e0 = op(op(e0, e1), e1)) & (e0 = op(op(e0, e0), e0)) & ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f9, e13, e12, e11, e10])).
fof(f10, plain, ((~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e10])).
fof(e10, plain, (sP0 <=> (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f11, plain, ((~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP1), inference(usedef, [], [e11])).
fof(e11, plain, (sP1 <=> (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f12, plain, ((~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP2), inference(usedef, [], [e12])).
fof(e12, plain, (sP2 <=> (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f13, plain, ((~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP3), inference(usedef, [], [e13])).
fof(e13, plain, (sP3 <=> (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f9, plain, ((e4 = op(op(e4, e4), e4)) & (e4 = op(op(e4, e3), e3)) & (e4 = op(op(e4, e2), e2)) & (e4 = op(op(e4, e1), e1)) & (e4 = op(op(e4, e0), e0)) & (e3 = op(op(e3, e4), e4)) & (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e2), e2)) & (e3 = op(op(e3, e1), e1)) & (e3 = op(op(e3, e0), e0)) & (e2 = op(op(e2, e4), e4)) & (e2 = op(op(e2, e3), e3)) & (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e1), e1)) & (e2 = op(op(e2, e0), e0)) & (e1 = op(op(e1, e4), e4)) & (e1 = op(op(e1, e3), e3)) & (e1 = op(op(e1, e2), e2)) & (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e0, e4), e4)) & (e0 = op(op(e0, e3), e3)) & (e0 = op(op(e0, e2), e2)) & (e0 = op(op(e0, e1), e1)) & (e0 = op(op(e0, e0), e0)) & ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0))))), inference(flattening, [], [f8])).
fof(f8, plain, ~ ~ ((e4 = op(op(e4, e4), e4)) & (e4 = op(op(e4, e3), e3)) & (e4 = op(op(e4, e2), e2)) & (e4 = op(op(e4, e1), e1)) & (e4 = op(op(e4, e0), e0)) & (e3 = op(op(e3, e4), e4)) & (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e2), e2)) & (e3 = op(op(e3, e1), e1)) & (e3 = op(op(e3, e0), e0)) & (e2 = op(op(e2, e4), e4)) & (e2 = op(op(e2, e3), e3)) & (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e1), e1)) & (e2 = op(op(e2, e0), e0)) & (e1 = op(op(e1, e4), e4)) & (e1 = op(op(e1, e3), e3)) & (e1 = op(op(e1, e2), e2)) & (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e0, e4), e4)) & (e0 = op(op(e0, e3), e3)) & (e0 = op(op(e0, e2), e2)) & (e0 = op(op(e0, e1), e1)) & (e0 = op(op(e0, e0), e0)) & ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0))))), inference(negated_conjecture, [], [f7])).
fof(f7, plain, ~ ~ ((e4 = op(op(e4, e4), e4)) & (e4 = op(op(e4, e3), e3)) & (e4 = op(op(e4, e2), e2)) & (e4 = op(op(e4, e1), e1)) & (e4 = op(op(e4, e0), e0)) & (e3 = op(op(e3, e4), e4)) & (e3 = op(op(e3, e3), e3)) & (e3 = op(op(e3, e2), e2)) & (e3 = op(op(e3, e1), e1)) & (e3 = op(op(e3, e0), e0)) & (e2 = op(op(e2, e4), e4)) & (e2 = op(op(e2, e3), e3)) & (e2 = op(op(e2, e2), e2)) & (e2 = op(op(e2, e1), e1)) & (e2 = op(op(e2, e0), e0)) & (e1 = op(op(e1, e4), e4)) & (e1 = op(op(e1, e3), e3)) & (e1 = op(op(e1, e2), e2)) & (e1 = op(op(e1, e1), e1)) & (e1 = op(op(e1, e0), e0)) & (e0 = op(op(e0, e4), e4)) & (e0 = op(op(e0, e3), e3)) & (e0 = op(op(e0, e2), e2)) & (e0 = op(op(e0, e1), e1)) & (e0 = op(op(e0, e0), e0)) & ((~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG174+1.p', co1)).
fof(f858, plain, spl4_33, inference(avatar_split_clause, [], [f231, f418])).
fof(f231, plain, (e2 = op(e3, e3)), inference(cnf_transformation, [], [f6])).