fof(f30403, plain, $false, inference(avatar_sat_refutation, [], [f1821, f1879, f1908, f2024, f2053, f2169, f2227, f2256, f2285, f2575, f2633, f2691, f2720, f2836, f2923, f2981, f3012, f3019, f3020, f3028, f3030, f3034, f3045, f3051, f3053, f3064, f3066, f3067, f3071, f3074, f3077, f3078, f3081, f3084, f3085, f3093, f3094, f3097, f3101, f3195, f3253, f3311, f3340, f3427, f3514, f3543, f3746, f3775, f3833, f3891, f3949, f4123, f4152, f4210, f4239, f4268, f4413, f4471, f4500, f4531, f4534, f4538, f4540, f4541, f4544, f4545, f4547, f4548, f4549, f4553, f4560, f4570, f4574, f4575, f4580, f4581, f4586, f4599, f4601, f4602, f4604, f4613, f4623, f4629, f4631, f4632, f4634, f4637, f4638, f4640, f4643, f5684, f5704, f5759, f5779, f5820, f8081, f8610, f8612, f8615, f8707, f8717, f8718, f8751, f8778, f8802, f8824, f8848, f8874, f8942, f8962, f8983, f9107, f9221, f9246, f9311, f9324, f9329, f9368, f9387, f9396, f9399, f9457, f9480, f9524, f9573, f9582, f9682, f9763, f9785, f9894, f9997, f10242, f10401, f10404, f10518, f10521, f10540, f10548, f10569, f10615, f10621, f10625, f10633, f10654, f10656, f10662, f10668, f10696, f10700, f10733, f10736, f10748, f10771, f10778, f10784, f10787, f10791, f10858, f10895, f10898, f10938, f10941, f10963, f11012, f11115, f11144, f11150, f11183, f11184, f11185, f11187, f11207, f11340, f11365, f11486, f11500, f11605, f11650, f11656, f11682, f11694, f11735, f11812, f11854, f11903, f11962, f11965, f11989, f12000, f12106, f12116, f12122, f12153, f12159, f12183, f12207, f12218, f12219, f12220, f12221, f12223, f12241, f12242, f12243, f12244, f12254, f12265, f12266, f12293, f12296, f12303, f12322, f12336, f12361, f12381, f12390, f12418, f12430, f12441, f12466, f12493, f12516, f12518, f12520, f12524, f12540, f12549, f12550, f12551, f12573, f12580, f12585, f12596, f12633, f12634, f12635, f12636, f12641, f12653, f12684, f12691, f12695, f12734, f12765, f12819, f12848, f12866, f12876, f12924, f12941, f12945, f12947, f12949, f12951, f12963, f13010, f13014, f13050, f13176, f13215, f13245, f13275, f13343, f13409, f13457, f13493, f13531, f13537, f13642, f13748, f13750, f13786, f13797, f13798, f13826, f13858, f13900, f14002, f14028, f14044, f14084, f14086, f14136, f14211, f14304, f14427, f14469, f14518, f14520, f14549, f14576, f14611, f14683, f14685, f14690, f14702, f14725, f14742, f14761, f14803, f14837, f14891, f14934, f14965, f14983, f15010, f15043, f15103, f15199, f15216, f15217, f15218, f15260, f15281, f15332, f15409, f15463, f15464, f15477, f15494, f15506, f15522, f15533, f15691, f15699, f15705, f15755, f15763, f15776, f15820, f15853, f15871, f15883, f15888, f15909, f15970, f15976, f16008, f16024, f16052, f16220, f16312, f16361, f16374, f16379, f16399, f16408, f16430, f16458, f16493, f16520, f16523, f16632, f16647, f16699, f16721, f16743, f16794, f16795, f16871, f16907, f16916, f16929, f16968, f16989, f17017, f17061, f17115, f17128, f17130, f17142, f17144, f17183, f17198, f17222, f17244, f17264, f17301, f17344, f17408, f17428, f17441, f17494, f17499, f17531, f17535, f17537, f17546, f17586, f17600, f17614, f17615, f17626, f17633, f17690, f17721, f17736, f17802, f17822, f17842, f17898, f17915, f17942, f17999, f18042, f18085, f18097, f18100, f18120, f18124, f18143, f18159, f18197, f18210, f18449, f18520, f18530, f18563, f18576, f18578, f18583, f18599, f18622, f18631, f18646, f18684, f18686, f18787, f18817, f18837, f18839, f18875, f18879, f18934, f18984, f19144, f19173, f19190, f19220, f19258, f19325, f19356, f19363, f19385, f19561, f19579, f19595, f19604, f19620, f19643, f19693, f19718, f19873, f19948, f19951, f20069, f20081, f20112, f20132, f20144, f20166, f20188, f20263, f20282, f20314, f20360, f20402, f20449, f20457, f20481, f20553, f20572, f20662, f20712, f20730, f20737, f20756, f20783, f20790, f20797, f20807, f20811, f20823, f20829, f20830, f20843, f20889, f20893, f20898, f20995, f21191, f21199, f21245, f21318, f21411, f21456, f21465, f21481, f21505, f21563, f21595, f21623, f21633, f21643, f21707, f21731, f21754, f21790, f21814, f21825, f21844, f21976, f22017, f22029, f22063, f22082, f22119, f22178, f22188, f22244, f22246, f22283, f22284, f22297, f22324, f22347, f22404, f22420, f22564, f22588, f22645, f22656, f22675, f22803, f22832, f22895, f22899, f22943, f22977, f23150, f23191, f23192, f23251, f23252, f23356, f23369, f23379, f23387, f23460, f23471, f23530, f23689, f23713, f23738, f23759, f23762, f23800, f23842, f23862, f23952, f23968, f24012, f24043, f24098, f24135, f24144, f24194, f24196, f24240, f24256, f24276, f24288, f24305, f24316, f24334, f24385, f24427, f24466, f24496, f24517, f24526, f24542, f24606, f24627, f24670, f24696, f24725, f24814, f24887, f24949, f25012, f25072, f25134, f25200, f25261, f25325, f25391, f25453, f25577, f25636, f25695, f25757, f25807, f25886, f25909, f25989, f26053, f26119, f26184, f26249, f26315, f26379, f26437, f26492, f26542, f26597, f26657, f26717, f26774, f26832, f26893, f26957, f27017, f27072, f27123, f27176, f27226, f27277, f27321, f27372, f27428, f27478, f27523, f27579, f27583, f27600, f27777, f27793, f27820, f27854, f27958, f28207, f28250, f28325, f28542, f28639, f28655, f28668, f28691, f28716, f30172, f30194, f30244, f30292, f30331, f30355, f30371, f30374, f30386])).
fof(f30386, plain, (~ spl42_388 | spl42_540 | ~ spl42_542), inference(avatar_contradiction_clause, [], [f30385])).
fof(f30385, plain, ($false | (~ spl42_388 | spl42_540 | ~ spl42_542)), inference(subsumption_resolution, [], [f30384, f3923])).
fof(f3923, plain, (~ (e20 = op2(e22, e26)) | spl42_540), inference(avatar_component_clause, [], [f3922])).
fof(f3922, plain, (spl42_540 <=> (e20 = op2(e22, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_540])])).
fof(f30384, plain, ((e20 = op2(e22, e26)) | (~ spl42_388 | ~ spl42_542)), inference(forward_demodulation, [], [f30375, f3932])).
fof(f3932, plain, ((e22 = op2(e22, e26)) | ~ spl42_542), inference(avatar_component_clause, [], [f3930])).
fof(f3930, plain, (spl42_542 <=> (e22 = op2(e22, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_542])])).
fof(f30375, plain, ((e20 = op2(op2(e22, e26), e26)) | ~ spl42_388), inference(backward_demodulation, [], [f460, f3294])).
fof(f3294, plain, ((e22 = op2(e26, e20)) | ~ spl42_388), inference(avatar_component_clause, [], [f3292])).
fof(f3292, plain, (spl42_388 <=> (e22 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_388])])).
fof(f460, plain, (e20 = op2(op2(op2(e26, e20), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ((e26 = op2(op2(op2(e26, e26), e26), e26)) & (e26 = op2(op2(op2(e25, e26), e25), e25)) & (e26 = op2(op2(op2(e24, e26), e24), e24)) & (e26 = op2(op2(op2(e23, e26), e23), e23)) & (e26 = op2(op2(op2(e22, e26), e22), e22)) & (e26 = op2(op2(op2(e21, e26), e21), e21)) & (e26 = op2(op2(op2(e20, e26), e20), e20)) & (e25 = op2(op2(op2(e26, e25), e26), e26)) & (e25 = op2(op2(op2(e25, e25), e25), e25)) & (e25 = op2(op2(op2(e24, e25), e24), e24)) & (e25 = op2(op2(op2(e23, e25), e23), e23)) & (e25 = op2(op2(op2(e22, e25), e22), e22)) & (e25 = op2(op2(op2(e21, e25), e21), e21)) & (e25 = op2(op2(op2(e20, e25), e20), e20)) & (e24 = op2(op2(op2(e26, e24), e26), e26)) & (e24 = op2(op2(op2(e25, e24), e25), e25)) & (e24 = op2(op2(op2(e24, e24), e24), e24)) & (e24 = op2(op2(op2(e23, e24), e23), e23)) & (e24 = op2(op2(op2(e22, e24), e22), e22)) & (e24 = op2(op2(op2(e21, e24), e21), e21)) & (e24 = op2(op2(op2(e20, e24), e20), e20)) & (e23 = op2(op2(op2(e26, e23), e26), e26)) & (e23 = op2(op2(op2(e25, e23), e25), e25)) & (e23 = op2(op2(op2(e24, e23), e24), e24)) & (e23 = op2(op2(op2(e23, e23), e23), e23)) & (e23 = op2(op2(op2(e22, e23), e22), e22)) & (e23 = op2(op2(op2(e21, e23), e21), e21)) & (e23 = op2(op2(op2(e20, e23), e20), e20)) & (e22 = op2(op2(op2(e26, e22), e26), e26)) & (e22 = op2(op2(op2(e25, e22), e25), e25)) & (e22 = op2(op2(op2(e24, e22), e24), e24)) & (e22 = op2(op2(op2(e23, e22), e23), e23)) & (e22 = op2(op2(op2(e22, e22), e22), e22)) & (e22 = op2(op2(op2(e21, e22), e21), e21)) & (e22 = op2(op2(op2(e20, e22), e20), e20)) & (e21 = op2(op2(op2(e26, e21), e26), e26)) & (e21 = op2(op2(op2(e25, e21), e25), e25)) & (e21 = op2(op2(op2(e24, e21), e24), e24)) & (e21 = op2(op2(op2(e23, e21), e23), e23)) & (e21 = op2(op2(op2(e22, e21), e22), e22)) & (e21 = op2(op2(op2(e21, e21), e21), e21)) & (e21 = op2(op2(op2(e20, e21), e20), e20)) & (e20 = op2(op2(op2(e26, e20), e26), e26)) & (e20 = op2(op2(op2(e25, e20), e25), e25)) & (e20 = op2(op2(op2(e24, e20), e24), e24)) & (e20 = op2(op2(op2(e23, e20), e23), e23)) & (e20 = op2(op2(op2(e22, e20), e22), e22)) & (e20 = op2(op2(op2(e21, e20), e21), e21)) & (e20 = op2(op2(op2(e20, e20), e20), e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax6)).
fof(f30374, plain, (~ spl42_397 | ~ spl42_443 | ~ spl42_688), inference(avatar_contradiction_clause, [], [f30373])).
fof(f30373, plain, ($false | (~ spl42_397 | ~ spl42_443 | ~ spl42_688)), inference(subsumption_resolution, [], [f30372, f1122])).
fof(f1122, plain, ~ (e21 = e26), inference(cnf_transformation, [], [f10])).
fof(f10, plain, (~ (e25 = e26) & ~ (e24 = e26) & ~ (e24 = e25) & ~ (e23 = e26) & ~ (e23 = e25) & ~ (e23 = e24) & ~ (e22 = e26) & ~ (e22 = e25) & ~ (e22 = e24) & ~ (e22 = e23) & ~ (e21 = e26) & ~ (e21 = e25) & ~ (e21 = e24) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e26) & ~ (e20 = e25) & ~ (e20 = e24) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax10)).
fof(f30372, plain, ((e21 = e26) | (~ spl42_397 | ~ spl42_443 | ~ spl42_688)), inference(forward_demodulation, [], [f30367, f3522])).
fof(f3522, plain, ((e21 = op2(e24, e26)) | ~ spl42_443), inference(avatar_component_clause, [], [f3520])).
fof(f3520, plain, (spl42_443 <=> (e21 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_443])])).
fof(f30367, plain, ((e26 = op2(e24, e26)) | (~ spl42_397 | ~ spl42_688)), inference(backward_demodulation, [], [f12687, f3331])).
fof(f3331, plain, ((e24 = op2(e25, e26)) | ~ spl42_397), inference(avatar_component_clause, [], [f3329])).
fof(f3329, plain, (spl42_397 <=> (e24 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_397])])).
fof(f12687, plain, ((e26 = op2(op2(e25, e26), e26)) | ~ spl42_688), inference(forward_demodulation, [], [f4770, f4817])).
fof(f4817, plain, ((e25 = h7(e15)) | ~ spl42_688), inference(avatar_component_clause, [], [f4816])).
fof(f4816, plain, (spl42_688 <=> (e25 = h7(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_688])])).
fof(f4770, plain, (e26 = op2(op2(h7(e15), e26), e26)), inference(backward_demodulation, [], [f502, f1246])).
fof(f1246, plain, (op2(e26, e26) = h7(e15)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ((op2(e26, e26) = h7(e15)) & (op2(e26, op2(e26, op2(op2(e26, e26), op2(e26, e26)))) = h7(e14)) & (op2(e26, op2(op2(e26, e26), op2(e26, e26))) = h7(e13)) & (op2(op2(op2(e26, e26), op2(e26, e26)), op2(e26, e26)) = h7(e12)) & (op2(op2(e26, e26), op2(e26, e26)) = h7(e11)) & (op2(e26, op2(e26, e26)) = h7(e10)) & (e26 = h7(e16))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax22)).
fof(f502, plain, (e26 = op2(op2(op2(e26, e26), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f30371, plain, (~ spl42_397 | ~ spl42_453), inference(avatar_contradiction_clause, [], [f30370])).
fof(f30370, plain, ($false | (~ spl42_397 | ~ spl42_453)), inference(subsumption_resolution, [], [f30369, f1131])).
fof(f1131, plain, ~ (e24 = e26), inference(cnf_transformation, [], [f10])).
fof(f30369, plain, ((e24 = e26) | (~ spl42_397 | ~ spl42_453)), inference(forward_demodulation, [], [f30368, f3563])).
fof(f3563, plain, ((e24 = op2(e24, e25)) | ~ spl42_453), inference(avatar_component_clause, [], [f3561])).
fof(f3561, plain, (spl42_453 <=> (e24 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_453])])).
fof(f30368, plain, ((e26 = op2(e24, e25)) | (~ spl42_397 | ~ spl42_453)), inference(forward_demodulation, [], [f30364, f3563])).
fof(f30364, plain, ((e26 = op2(op2(e24, e25), e25)) | ~ spl42_397), inference(backward_demodulation, [], [f501, f3331])).
fof(f501, plain, (e26 = op2(op2(op2(e25, e26), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f30355, plain, (~ spl42_423 | ~ spl42_426), inference(avatar_contradiction_clause, [], [f30354])).
fof(f30354, plain, ($false | (~ spl42_423 | ~ spl42_426)), inference(subsumption_resolution, [], [f30353, f1125])).
fof(f1125, plain, ~ (e22 = e25), inference(cnf_transformation, [], [f10])).
fof(f30353, plain, ((e22 = e25) | (~ spl42_423 | ~ spl42_426)), inference(backward_demodulation, [], [f3451, f3439])).
fof(f3439, plain, ((e22 = op2(e25, e22)) | ~ spl42_423), inference(avatar_component_clause, [], [f3437])).
fof(f3437, plain, (spl42_423 <=> (e22 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_423])])).
fof(f3451, plain, ((e25 = op2(e25, e22)) | ~ spl42_426), inference(avatar_component_clause, [], [f3449])).
fof(f3449, plain, (spl42_426 <=> (e25 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_426])])).
fof(f30331, plain, (~ spl42_491 | spl42_639 | ~ spl42_698), inference(avatar_contradiction_clause, [], [f30330])).
fof(f30330, plain, ($false | (~ spl42_491 | spl42_639 | ~ spl42_698)), inference(subsumption_resolution, [], [f30327, f4333])).
fof(f4333, plain, (~ (e21 = op2(e20, e26)) | spl42_639), inference(avatar_component_clause, [], [f4332])).
fof(f4332, plain, (spl42_639 <=> (e21 = op2(e20, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_639])])).
fof(f30327, plain, ((e21 = op2(e20, e26)) | (~ spl42_491 | ~ spl42_698)), inference(backward_demodulation, [], [f15201, f3721])).
fof(f3721, plain, ((e20 = op2(e23, e26)) | ~ spl42_491), inference(avatar_component_clause, [], [f3719])).
fof(f3719, plain, (spl42_491 <=> (e20 = op2(e23, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_491])])).
fof(f15201, plain, ((e21 = op2(op2(e23, e26), e26)) | ~ spl42_698), inference(backward_demodulation, [], [f4790, f4873])).
fof(f4873, plain, ((e23 = h7(e13)) | ~ spl42_698), inference(avatar_component_clause, [], [f4872])).
fof(f4872, plain, (spl42_698 <=> (e23 = h7(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_698])])).
fof(f4790, plain, (e21 = op2(op2(h7(e13), e26), e26)), inference(backward_demodulation, [], [f467, f4789])).
fof(f4789, plain, (op2(e26, e21) = h7(e13)), inference(forward_demodulation, [], [f4788, f4784])).
fof(f4784, plain, (e21 = op2(h7(e15), h7(e15))), inference(backward_demodulation, [], [f1193, f1246])).
fof(f1193, plain, (e21 = op2(op2(e26, e26), op2(e26, e26))), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((e25 = op2(e26, e26)) & (e24 = op2(e26, op2(e26, op2(op2(e26, e26), op2(e26, e26))))) & (e23 = op2(e26, op2(op2(e26, e26), op2(e26, e26)))) & (e22 = op2(op2(op2(e26, e26), op2(e26, e26)), op2(e26, e26))) & (e21 = op2(op2(e26, e26), op2(e26, e26))) & (e20 = op2(e26, op2(e26, e26)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax15)).
fof(f4788, plain, (h7(e13) = op2(e26, op2(h7(e15), h7(e15)))), inference(forward_demodulation, [], [f1244, f1246])).
fof(f1244, plain, (op2(e26, op2(op2(e26, e26), op2(e26, e26))) = h7(e13)), inference(cnf_transformation, [], [f22])).
fof(f467, plain, (e21 = op2(op2(op2(e26, e21), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f30292, plain, (spl42_483 | ~ spl42_530 | ~ spl42_592), inference(avatar_split_clause, [], [f30291, f4137, f3880, f3685])).
fof(f3685, plain, (spl42_483 <=> (e26 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_483])])).
fof(f3880, plain, (spl42_530 <=> (e24 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_530])])).
fof(f4137, plain, (spl42_592 <=> (e23 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_592])])).
fof(f30291, plain, ((e26 = op2(e24, e21)) | (~ spl42_530 | ~ spl42_592)), inference(backward_demodulation, [], [f29579, f3882])).
fof(f3882, plain, ((e24 = op2(e23, e21)) | ~ spl42_530), inference(avatar_component_clause, [], [f3880])).
fof(f29579, plain, ((e26 = op2(op2(e23, e21), e21)) | ~ spl42_592), inference(forward_demodulation, [], [f497, f4139])).
fof(f4139, plain, ((e23 = op2(e21, e26)) | ~ spl42_592), inference(avatar_component_clause, [], [f4137])).
fof(f497, plain, (e26 = op2(op2(op2(e21, e26), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f30244, plain, (spl42_364 | ~ spl42_443 | ~ spl42_609), inference(avatar_split_clause, [], [f30240, f4207, f3520, f3192])).
fof(f3192, plain, (spl42_364 <=> (e26 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_364])])).
fof(f4207, plain, (spl42_609 <=> (e26 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_609])])).
fof(f30240, plain, ((e26 = op2(e26, e24)) | (~ spl42_443 | ~ spl42_609)), inference(backward_demodulation, [], [f24136, f4209])).
fof(f4209, plain, ((e26 = op2(e21, e24)) | ~ spl42_609), inference(avatar_component_clause, [], [f4207])).
fof(f24136, plain, ((e26 = op2(op2(e21, e24), e24)) | ~ spl42_443), inference(backward_demodulation, [], [f500, f3522])).
fof(f500, plain, (e26 = op2(op2(op2(e24, e26), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f30194, plain, (~ spl42_644 | ~ spl42_672), inference(avatar_split_clause, [], [f30192, f4468, f4352])).
fof(f4352, plain, (spl42_644 <=> (e26 = op2(e20, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_644])])).
fof(f4468, plain, (spl42_672 <=> (e26 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_672])])).
fof(f30192, plain, (~ (e26 = op2(e20, e26)) | ~ spl42_672), inference(backward_demodulation, [], [f958, f4470])).
fof(f4470, plain, ((e26 = op2(e20, e22)) | ~ spl42_672), inference(avatar_component_clause, [], [f4468])).
fof(f958, plain, ~ (op2(e20, e22) = op2(e20, e26)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (op2(e26, e25) = op2(e26, e26)) & ~ (op2(e26, e24) = op2(e26, e26)) & ~ (op2(e26, e24) = op2(e26, e25)) & ~ (op2(e26, e23) = op2(e26, e26)) & ~ (op2(e26, e23) = op2(e26, e25)) & ~ (op2(e26, e23) = op2(e26, e24)) & ~ (op2(e26, e22) = op2(e26, e26)) & ~ (op2(e26, e22) = op2(e26, e25)) & ~ (op2(e26, e22) = op2(e26, e24)) & ~ (op2(e26, e22) = op2(e26, e23)) & ~ (op2(e26, e21) = op2(e26, e26)) & ~ (op2(e26, e21) = op2(e26, e25)) & ~ (op2(e26, e21) = op2(e26, e24)) & ~ (op2(e26, e21) = op2(e26, e23)) & ~ (op2(e26, e21) = op2(e26, e22)) & ~ (op2(e26, e20) = op2(e26, e26)) & ~ (op2(e26, e20) = op2(e26, e25)) & ~ (op2(e26, e20) = op2(e26, e24)) & ~ (op2(e26, e20) = op2(e26, e23)) & ~ (op2(e26, e20) = op2(e26, e22)) & ~ (op2(e26, e20) = op2(e26, e21)) & ~ (op2(e25, e25) = op2(e25, e26)) & ~ (op2(e25, e24) = op2(e25, e26)) & ~ (op2(e25, e24) = op2(e25, e25)) & ~ (op2(e25, e23) = op2(e25, e26)) & ~ (op2(e25, e23) = op2(e25, e25)) & ~ (op2(e25, e23) = op2(e25, e24)) & ~ (op2(e25, e22) = op2(e25, e26)) & ~ (op2(e25, e22) = op2(e25, e25)) & ~ (op2(e25, e22) = op2(e25, e24)) & ~ (op2(e25, e22) = op2(e25, e23)) & ~ (op2(e25, e21) = op2(e25, e26)) & ~ (op2(e25, e21) = op2(e25, e25)) & ~ (op2(e25, e21) = op2(e25, e24)) & ~ (op2(e25, e21) = op2(e25, e23)) & ~ (op2(e25, e21) = op2(e25, e22)) & ~ (op2(e25, e20) = op2(e25, e26)) & ~ (op2(e25, e20) = op2(e25, e25)) & ~ (op2(e25, e20) = op2(e25, e24)) & ~ (op2(e25, e20) = op2(e25, e23)) & ~ (op2(e25, e20) = op2(e25, e22)) & ~ (op2(e25, e20) = op2(e25, e21)) & ~ (op2(e24, e25) = op2(e24, e26)) & ~ (op2(e24, e24) = op2(e24, e26)) & ~ (op2(e24, e24) = op2(e24, e25)) & ~ (op2(e24, e23) = op2(e24, e26)) & ~ (op2(e24, e23) = op2(e24, e25)) & ~ (op2(e24, e23) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e26)) & ~ (op2(e24, e22) = op2(e24, e25)) & ~ (op2(e24, e22) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e26)) & ~ (op2(e24, e21) = op2(e24, e25)) & ~ (op2(e24, e21) = op2(e24, e24)) & ~ (op2(e24, e21) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e26)) & ~ (op2(e24, e20) = op2(e24, e25)) & ~ (op2(e24, e20) = op2(e24, e24)) & ~ (op2(e24, e20) = op2(e24, e23)) & ~ (op2(e24, e20) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e21)) & ~ (op2(e23, e25) = op2(e23, e26)) & ~ (op2(e23, e24) = op2(e23, e26)) & ~ (op2(e23, e24) = op2(e23, e25)) & ~ (op2(e23, e23) = op2(e23, e26)) & ~ (op2(e23, e23) = op2(e23, e25)) & ~ (op2(e23, e23) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e26)) & ~ (op2(e23, e22) = op2(e23, e25)) & ~ (op2(e23, e22) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e26)) & ~ (op2(e23, e21) = op2(e23, e25)) & ~ (op2(e23, e21) = op2(e23, e24)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e26)) & ~ (op2(e23, e20) = op2(e23, e25)) & ~ (op2(e23, e20) = op2(e23, e24)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e25) = op2(e22, e26)) & ~ (op2(e22, e24) = op2(e22, e26)) & ~ (op2(e22, e24) = op2(e22, e25)) & ~ (op2(e22, e23) = op2(e22, e26)) & ~ (op2(e22, e23) = op2(e22, e25)) & ~ (op2(e22, e23) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e26)) & ~ (op2(e22, e22) = op2(e22, e25)) & ~ (op2(e22, e22) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e26)) & ~ (op2(e22, e21) = op2(e22, e25)) & ~ (op2(e22, e21) = op2(e22, e24)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e26)) & ~ (op2(e22, e20) = op2(e22, e25)) & ~ (op2(e22, e20) = op2(e22, e24)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e25) = op2(e21, e26)) & ~ (op2(e21, e24) = op2(e21, e26)) & ~ (op2(e21, e24) = op2(e21, e25)) & ~ (op2(e21, e23) = op2(e21, e26)) & ~ (op2(e21, e23) = op2(e21, e25)) & ~ (op2(e21, e23) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e26)) & ~ (op2(e21, e22) = op2(e21, e25)) & ~ (op2(e21, e22) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e26)) & ~ (op2(e21, e21) = op2(e21, e25)) & ~ (op2(e21, e21) = op2(e21, e24)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e26)) & ~ (op2(e21, e20) = op2(e21, e25)) & ~ (op2(e21, e20) = op2(e21, e24)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e25) = op2(e20, e26)) & ~ (op2(e20, e24) = op2(e20, e26)) & ~ (op2(e20, e24) = op2(e20, e25)) & ~ (op2(e20, e23) = op2(e20, e26)) & ~ (op2(e20, e23) = op2(e20, e25)) & ~ (op2(e20, e23) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e26)) & ~ (op2(e20, e22) = op2(e20, e25)) & ~ (op2(e20, e22) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e26)) & ~ (op2(e20, e21) = op2(e20, e25)) & ~ (op2(e20, e21) = op2(e20, e24)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e26)) & ~ (op2(e20, e20) = op2(e20, e25)) & ~ (op2(e20, e20) = op2(e20, e24)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e25, e26) = op2(e26, e26)) & ~ (op2(e24, e26) = op2(e26, e26)) & ~ (op2(e24, e26) = op2(e25, e26)) & ~ (op2(e23, e26) = op2(e26, e26)) & ~ (op2(e23, e26) = op2(e25, e26)) & ~ (op2(e23, e26) = op2(e24, e26)) & ~ (op2(e22, e26) = op2(e26, e26)) & ~ (op2(e22, e26) = op2(e25, e26)) & ~ (op2(e22, e26) = op2(e24, e26)) & ~ (op2(e22, e26) = op2(e23, e26)) & ~ (op2(e21, e26) = op2(e26, e26)) & ~ (op2(e21, e26) = op2(e25, e26)) & ~ (op2(e21, e26) = op2(e24, e26)) & ~ (op2(e21, e26) = op2(e23, e26)) & ~ (op2(e21, e26) = op2(e22, e26)) & ~ (op2(e20, e26) = op2(e26, e26)) & ~ (op2(e20, e26) = op2(e25, e26)) & ~ (op2(e20, e26) = op2(e24, e26)) & ~ (op2(e20, e26) = op2(e23, e26)) & ~ (op2(e20, e26) = op2(e22, e26)) & ~ (op2(e20, e26) = op2(e21, e26)) & ~ (op2(e25, e25) = op2(e26, e25)) & ~ (op2(e24, e25) = op2(e26, e25)) & ~ (op2(e24, e25) = op2(e25, e25)) & ~ (op2(e23, e25) = op2(e26, e25)) & ~ (op2(e23, e25) = op2(e25, e25)) & ~ (op2(e23, e25) = op2(e24, e25)) & ~ (op2(e22, e25) = op2(e26, e25)) & ~ (op2(e22, e25) = op2(e25, e25)) & ~ (op2(e22, e25) = op2(e24, e25)) & ~ (op2(e22, e25) = op2(e23, e25)) & ~ (op2(e21, e25) = op2(e26, e25)) & ~ (op2(e21, e25) = op2(e25, e25)) & ~ (op2(e21, e25) = op2(e24, e25)) & ~ (op2(e21, e25) = op2(e23, e25)) & ~ (op2(e21, e25) = op2(e22, e25)) & ~ (op2(e20, e25) = op2(e26, e25)) & ~ (op2(e20, e25) = op2(e25, e25)) & ~ (op2(e20, e25) = op2(e24, e25)) & ~ (op2(e20, e25) = op2(e23, e25)) & ~ (op2(e20, e25) = op2(e22, e25)) & ~ (op2(e20, e25) = op2(e21, e25)) & ~ (op2(e25, e24) = op2(e26, e24)) & ~ (op2(e24, e24) = op2(e26, e24)) & ~ (op2(e24, e24) = op2(e25, e24)) & ~ (op2(e23, e24) = op2(e26, e24)) & ~ (op2(e23, e24) = op2(e25, e24)) & ~ (op2(e23, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e26, e24)) & ~ (op2(e22, e24) = op2(e25, e24)) & ~ (op2(e22, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e26, e24)) & ~ (op2(e21, e24) = op2(e25, e24)) & ~ (op2(e21, e24) = op2(e24, e24)) & ~ (op2(e21, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e26, e24)) & ~ (op2(e20, e24) = op2(e25, e24)) & ~ (op2(e20, e24) = op2(e24, e24)) & ~ (op2(e20, e24) = op2(e23, e24)) & ~ (op2(e20, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e21, e24)) & ~ (op2(e25, e23) = op2(e26, e23)) & ~ (op2(e24, e23) = op2(e26, e23)) & ~ (op2(e24, e23) = op2(e25, e23)) & ~ (op2(e23, e23) = op2(e26, e23)) & ~ (op2(e23, e23) = op2(e25, e23)) & ~ (op2(e23, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e26, e23)) & ~ (op2(e22, e23) = op2(e25, e23)) & ~ (op2(e22, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e26, e23)) & ~ (op2(e21, e23) = op2(e25, e23)) & ~ (op2(e21, e23) = op2(e24, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e26, e23)) & ~ (op2(e20, e23) = op2(e25, e23)) & ~ (op2(e20, e23) = op2(e24, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e25, e22) = op2(e26, e22)) & ~ (op2(e24, e22) = op2(e26, e22)) & ~ (op2(e24, e22) = op2(e25, e22)) & ~ (op2(e23, e22) = op2(e26, e22)) & ~ (op2(e23, e22) = op2(e25, e22)) & ~ (op2(e23, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e26, e22)) & ~ (op2(e22, e22) = op2(e25, e22)) & ~ (op2(e22, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e26, e22)) & ~ (op2(e21, e22) = op2(e25, e22)) & ~ (op2(e21, e22) = op2(e24, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e26, e22)) & ~ (op2(e20, e22) = op2(e25, e22)) & ~ (op2(e20, e22) = op2(e24, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e25, e21) = op2(e26, e21)) & ~ (op2(e24, e21) = op2(e26, e21)) & ~ (op2(e24, e21) = op2(e25, e21)) & ~ (op2(e23, e21) = op2(e26, e21)) & ~ (op2(e23, e21) = op2(e25, e21)) & ~ (op2(e23, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e26, e21)) & ~ (op2(e22, e21) = op2(e25, e21)) & ~ (op2(e22, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e26, e21)) & ~ (op2(e21, e21) = op2(e25, e21)) & ~ (op2(e21, e21) = op2(e24, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e26, e21)) & ~ (op2(e20, e21) = op2(e25, e21)) & ~ (op2(e20, e21) = op2(e24, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e25, e20) = op2(e26, e20)) & ~ (op2(e24, e20) = op2(e26, e20)) & ~ (op2(e24, e20) = op2(e25, e20)) & ~ (op2(e23, e20) = op2(e26, e20)) & ~ (op2(e23, e20) = op2(e25, e20)) & ~ (op2(e23, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e26, e20)) & ~ (op2(e22, e20) = op2(e25, e20)) & ~ (op2(e22, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e26, e20)) & ~ (op2(e21, e20) = op2(e25, e20)) & ~ (op2(e21, e20) = op2(e24, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e26, e20)) & ~ (op2(e20, e20) = op2(e25, e20)) & ~ (op2(e20, e20) = op2(e24, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax8)).
fof(f30172, plain, (~ spl42_735 | ~ spl42_742), inference(avatar_contradiction_clause, [], [f30171])).
fof(f30171, plain, ($false | (~ spl42_735 | ~ spl42_742)), inference(subsumption_resolution, [], [f30170, f1112])).
fof(f1112, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f10])).
fof(f30170, plain, ((e20 = e21) | (~ spl42_735 | ~ spl42_742)), inference(forward_demodulation, [], [f5074, f5109])).
fof(f5109, plain, ((e20 = h6(e13)) | ~ spl42_742), inference(avatar_component_clause, [], [f5108])).
fof(f5108, plain, (spl42_742 <=> (e20 = h6(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_742])])).
fof(f5074, plain, ((e21 = h6(e13)) | ~ spl42_735), inference(avatar_component_clause, [], [f5073])).
fof(f5073, plain, (spl42_735 <=> (e21 = h6(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_735])])).
fof(f28716, plain, (~ spl42_563 | ~ spl42_542), inference(avatar_split_clause, [], [f27794, f3930, f4017])).
fof(f4017, plain, (spl42_563 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_563])])).
fof(f27794, plain, (~ (e22 = op2(e22, e23)) | ~ spl42_542), inference(forward_demodulation, [], [f1003, f3932])).
fof(f1003, plain, ~ (op2(e22, e23) = op2(e22, e26)), inference(cnf_transformation, [], [f8])).
fof(f28691, plain, (spl42_468 | ~ spl42_369 | ~ spl42_504), inference(avatar_split_clause, [], [f28599, f3772, f3213, f3623])).
fof(f3623, plain, (spl42_468 <=> (e25 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_468])])).
fof(f3213, plain, (spl42_369 <=> (e24 = op2(e26, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_369])])).
fof(f3772, plain, (spl42_504 <=> (e26 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_504])])).
fof(f28599, plain, ((e25 = op2(e24, e23)) | (~ spl42_369 | ~ spl42_504)), inference(forward_demodulation, [], [f28596, f3215])).
fof(f3215, plain, ((e24 = op2(e26, e23)) | ~ spl42_369), inference(avatar_component_clause, [], [f3213])).
fof(f28596, plain, ((e25 = op2(op2(e26, e23), e23)) | ~ spl42_504), inference(backward_demodulation, [], [f492, f3774])).
fof(f3774, plain, ((e26 = op2(e23, e25)) | ~ spl42_504), inference(avatar_component_clause, [], [f3772])).
fof(f492, plain, (e25 = op2(op2(op2(e23, e25), e23), e23)), inference(cnf_transformation, [], [f6])).
fof(f28668, plain, (~ spl42_401 | ~ spl42_409 | ~ spl42_552), inference(avatar_contradiction_clause, [], [f28667])).
fof(f28667, plain, ($false | (~ spl42_401 | ~ spl42_409 | ~ spl42_552)), inference(subsumption_resolution, [], [f28666, f1120])).
fof(f1120, plain, ~ (e21 = e24), inference(cnf_transformation, [], [f10])).
fof(f28666, plain, ((e21 = e24) | (~ spl42_401 | ~ spl42_409 | ~ spl42_552)), inference(forward_demodulation, [], [f28665, f3348])).
fof(f3348, plain, ((e21 = op2(e25, e25)) | ~ spl42_401), inference(avatar_component_clause, [], [f3346])).
fof(f3346, plain, (spl42_401 <=> (e21 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_401])])).
fof(f28665, plain, ((e24 = op2(e25, e25)) | (~ spl42_409 | ~ spl42_552)), inference(forward_demodulation, [], [f28659, f3973])).
fof(f3973, plain, ((e25 = op2(e22, e25)) | ~ spl42_552), inference(avatar_component_clause, [], [f3971])).
fof(f3971, plain, (spl42_552 <=> (e25 = op2(e22, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_552])])).
fof(f28659, plain, ((e24 = op2(op2(e22, e25), e25)) | ~ spl42_409), inference(backward_demodulation, [], [f487, f3381])).
fof(f3381, plain, ((e22 = op2(e25, e24)) | ~ spl42_409), inference(avatar_component_clause, [], [f3379])).
fof(f3379, plain, (spl42_409 <=> (e22 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_409])])).
fof(f487, plain, (e24 = op2(op2(op2(e25, e24), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f28655, plain, (~ spl42_413 | ~ spl42_420), inference(avatar_split_clause, [], [f28651, f3424, f3395])).
fof(f3395, plain, (spl42_413 <=> (e26 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_413])])).
fof(f3424, plain, (spl42_420 <=> (e26 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_420])])).
fof(f28651, plain, (~ (e26 = op2(e25, e24)) | ~ spl42_420), inference(backward_demodulation, [], [f1064, f3426])).
fof(f3426, plain, ((e26 = op2(e25, e23)) | ~ spl42_420), inference(avatar_component_clause, [], [f3424])).
fof(f1064, plain, ~ (op2(e25, e23) = op2(e25, e24)), inference(cnf_transformation, [], [f8])).
fof(f28639, plain, (spl42_412 | ~ spl42_453 | ~ spl42_461), inference(avatar_contradiction_clause, [], [f28638])).
fof(f28638, plain, ($false | (spl42_412 | ~ spl42_453 | ~ spl42_461)), inference(subsumption_resolution, [], [f28637, f3392])).
fof(f3392, plain, (~ (e25 = op2(e25, e24)) | spl42_412), inference(avatar_component_clause, [], [f3391])).
fof(f3391, plain, (spl42_412 <=> (e25 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_412])])).
fof(f28637, plain, ((e25 = op2(e25, e24)) | (~ spl42_453 | ~ spl42_461)), inference(forward_demodulation, [], [f28633, f3596])).
fof(f3596, plain, ((e25 = op2(e24, e24)) | ~ spl42_461), inference(avatar_component_clause, [], [f3594])).
fof(f3594, plain, (spl42_461 <=> (e25 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_461])])).
fof(f28633, plain, ((e25 = op2(op2(e24, e24), e24)) | ~ spl42_453), inference(backward_demodulation, [], [f493, f3563])).
fof(f493, plain, (e25 = op2(op2(op2(e24, e25), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f28542, plain, (spl42_462 | ~ spl42_443 | ~ spl42_607), inference(avatar_split_clause, [], [f28541, f4199, f3520, f3598])).
fof(f3598, plain, (spl42_462 <=> (e26 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_462])])).
fof(f4199, plain, (spl42_607 <=> (e24 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_607])])).
fof(f28541, plain, ((e26 = op2(e24, e24)) | (~ spl42_443 | ~ spl42_607)), inference(backward_demodulation, [], [f24136, f4201])).
fof(f4201, plain, ((e24 = op2(e21, e24)) | ~ spl42_607), inference(avatar_component_clause, [], [f4199])).
fof(f28325, plain, (~ spl42_413 | ~ spl42_1062), inference(avatar_split_clause, [], [f28303, f6974, f3395])).
fof(f6974, plain, (spl42_1062 <=> (e26 = h5(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_1062])])).
fof(f28303, plain, (~ (e26 = op2(e25, e24)) | ~ spl42_1062), inference(backward_demodulation, [], [f4733, f6975])).
fof(f6975, plain, ((e26 = h5(e15)) | ~ spl42_1062), inference(avatar_component_clause, [], [f6974])).
fof(f4733, plain, ~ (op2(e25, e24) = h5(e15)), inference(backward_demodulation, [], [f899, f1232])).
fof(f1232, plain, (op2(e24, e24) = h5(e15)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ((op2(e24, e24) = h5(e15)) & (h5(e14) = op2(e24, op2(e24, op2(op2(e24, e24), op2(e24, e24))))) & (h5(e13) = op2(e24, op2(op2(e24, e24), op2(e24, e24)))) & (h5(e12) = op2(op2(op2(e24, e24), op2(e24, e24)), op2(e24, e24))) & (h5(e11) = op2(op2(e24, e24), op2(e24, e24))) & (h5(e10) = op2(e24, op2(e24, e24))) & (e24 = h5(e16))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax20)).
fof(f899, plain, ~ (op2(e24, e24) = op2(e25, e24)), inference(cnf_transformation, [], [f8])).
fof(f28250, plain, (~ spl42_567 | ~ spl42_1172), inference(avatar_split_clause, [], [f28216, f7670, f4033])).
fof(f4033, plain, (spl42_567 <=> (e26 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_567])])).
fof(f7670, plain, (spl42_1172 <=> (e26 = h3(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_1172])])).
fof(f28216, plain, (~ (e26 = op2(e22, e23)) | ~ spl42_1172), inference(backward_demodulation, [], [f4695, f7671])).
fof(f7671, plain, ((e26 = h3(e15)) | ~ spl42_1172), inference(avatar_component_clause, [], [f7670])).
fof(f4695, plain, ~ (op2(e22, e23) = h3(e15)), inference(backward_demodulation, [], [f997, f1218])).
fof(f1218, plain, (op2(e22, e22) = h3(e15)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ((op2(e22, e22) = h3(e15)) & (h3(e14) = op2(e22, op2(e22, op2(op2(e22, e22), op2(e22, e22))))) & (h3(e13) = op2(e22, op2(op2(e22, e22), op2(e22, e22)))) & (h3(e12) = op2(op2(op2(e22, e22), op2(e22, e22)), op2(e22, e22))) & (h3(e11) = op2(op2(e22, e22), op2(e22, e22))) & (h3(e10) = op2(e22, op2(e22, e22))) & (e22 = h3(e16))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax18)).
fof(f997, plain, ~ (op2(e22, e22) = op2(e22, e23)), inference(cnf_transformation, [], [f8])).
fof(f28207, plain, (~ spl42_860 | ~ spl42_306 | ~ spl42_495 | ~ spl42_871 | spl42_1198 | ~ spl42_1227), inference(avatar_split_clause, [], [f28206, f8018, f7902, f5756, f3735, f2854, f5701])).
fof(f5701, plain, (spl42_860 <=> (e24 = h2(e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_860])])).
fof(f2854, plain, (spl42_306 <=> (e14 = op1(e10, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_306])])).
fof(f3735, plain, (spl42_495 <=> (e24 = op2(e23, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_495])])).
fof(f5756, plain, (spl42_871 <=> (e23 = h2(e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_871])])).
fof(f7902, plain, (spl42_1198 <=> (h2(op1(e10, e15)) = op2(h2(e10), h2(e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1198])])).
fof(f8018, plain, (spl42_1227 <=> (e26 = h2(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_1227])])).
fof(f28206, plain, (~ (e24 = h2(e14)) | (~ spl42_306 | ~ spl42_495 | ~ spl42_871 | spl42_1198 | ~ spl42_1227)), inference(forward_demodulation, [], [f28205, f3737])).
fof(f3737, plain, ((e24 = op2(e23, e26)) | ~ spl42_495), inference(avatar_component_clause, [], [f3735])).
fof(f28205, plain, (~ (op2(e23, e26) = h2(e14)) | (~ spl42_306 | ~ spl42_871 | spl42_1198 | ~ spl42_1227)), inference(forward_demodulation, [], [f28204, f2856])).
fof(f2856, plain, ((e14 = op1(e10, e15)) | ~ spl42_306), inference(avatar_component_clause, [], [f2854])).
fof(f28204, plain, (~ (op2(e23, e26) = h2(op1(e10, e15))) | (~ spl42_871 | spl42_1198 | ~ spl42_1227)), inference(forward_demodulation, [], [f28203, f5757])).
fof(f5757, plain, ((e23 = h2(e10)) | ~ spl42_871), inference(avatar_component_clause, [], [f5756])).
fof(f28203, plain, (~ (h2(op1(e10, e15)) = op2(h2(e10), e26)) | (spl42_1198 | ~ spl42_1227)), inference(forward_demodulation, [], [f7904, f8019])).
fof(f8019, plain, ((e26 = h2(e15)) | ~ spl42_1227), inference(avatar_component_clause, [], [f8018])).
fof(f7904, plain, (~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | spl42_1198), inference(avatar_component_clause, [], [f7902])).
fof(f27958, plain, (~ spl42_803 | ~ spl42_611), inference(avatar_split_clause, [], [f27957, f4216, f5415])).
fof(f5415, plain, (spl42_803 <=> (e21 = h4(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_803])])).
fof(f4216, plain, (spl42_611 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_611])])).
fof(f27957, plain, (~ (e21 = h4(e15)) | ~ spl42_611), inference(forward_demodulation, [], [f4709, f4218])).
fof(f4218, plain, ((e21 = op2(e21, e23)) | ~ spl42_611), inference(avatar_component_clause, [], [f4216])).
fof(f4709, plain, ~ (op2(e21, e23) = h4(e15)), inference(backward_demodulation, [], [f867, f1225])).
fof(f1225, plain, (op2(e23, e23) = h4(e15)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ((op2(e23, e23) = h4(e15)) & (h4(e14) = op2(e23, op2(e23, op2(op2(e23, e23), op2(e23, e23))))) & (h4(e13) = op2(e23, op2(op2(e23, e23), op2(e23, e23)))) & (h4(e12) = op2(op2(op2(e23, e23), op2(e23, e23)), op2(e23, e23))) & (h4(e11) = op2(op2(e23, e23), op2(e23, e23))) & (h4(e10) = op2(e23, op2(e23, e23))) & (e23 = h4(e16))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax19)).
fof(f867, plain, ~ (op2(e21, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f27854, plain, (~ spl42_674 | ~ spl42_667), inference(avatar_split_clause, [], [f24551, f4448, f4477])).
fof(f4477, plain, (spl42_674 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_674])])).
fof(f4448, plain, (spl42_667 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_667])])).
fof(f24551, plain, (~ (e21 = op2(e20, e21)) | ~ spl42_667), inference(forward_demodulation, [], [f950, f4450])).
fof(f4450, plain, ((e21 = op2(e20, e22)) | ~ spl42_667), inference(avatar_component_clause, [], [f4448])).
fof(f950, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f8])).
fof(f27820, plain, (~ spl42_588 | ~ spl42_392), inference(avatar_split_clause, [], [f27587, f3308, f4120])).
fof(f4120, plain, (spl42_588 <=> (e26 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_588])])).
fof(f3308, plain, (spl42_392 <=> (e26 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_392])])).
fof(f27587, plain, (~ (e26 = op2(e22, e20)) | ~ spl42_392), inference(forward_demodulation, [], [f811, f3310])).
fof(f3310, plain, ((e26 = op2(e26, e20)) | ~ spl42_392), inference(avatar_component_clause, [], [f3308])).
fof(f811, plain, ~ (op2(e22, e20) = op2(e26, e20)), inference(cnf_transformation, [], [f8])).
fof(f27793, plain, (spl42_567 | ~ spl42_465 | ~ spl42_495), inference(avatar_split_clause, [], [f24097, f3735, f3611, f4033])).
fof(f3611, plain, (spl42_465 <=> (e22 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_465])])).
fof(f24097, plain, ((e26 = op2(e22, e23)) | (~ spl42_465 | ~ spl42_495)), inference(forward_demodulation, [], [f24091, f3613])).
fof(f3613, plain, ((e22 = op2(e24, e23)) | ~ spl42_465), inference(avatar_component_clause, [], [f3611])).
fof(f24091, plain, ((e26 = op2(op2(e24, e23), e23)) | ~ spl42_495), inference(backward_demodulation, [], [f499, f3737])).
fof(f499, plain, (e26 = op2(op2(op2(e23, e26), e23), e23)), inference(cnf_transformation, [], [f6])).
fof(f27777, plain, (~ spl42_524 | ~ spl42_426), inference(avatar_split_clause, [], [f20759, f3449, f3855])).
fof(f3855, plain, (spl42_524 <=> (e25 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_524])])).
fof(f20759, plain, (~ (e25 = op2(e23, e22)) | ~ spl42_426), inference(forward_demodulation, [], [f855, f3451])).
fof(f855, plain, ~ (op2(e23, e22) = op2(e25, e22)), inference(cnf_transformation, [], [f8])).
fof(f27600, plain, (~ spl42_339 | ~ spl42_512 | ~ spl42_871 | ~ spl42_883 | spl42_1193), inference(avatar_contradiction_clause, [], [f27599])).
fof(f27599, plain, ($false | (~ spl42_339 | ~ spl42_512 | ~ spl42_871 | ~ spl42_883 | spl42_1193)), inference(subsumption_resolution, [], [f27598, f3808])).
fof(f3808, plain, ((e20 = op2(e23, e23)) | ~ spl42_512), inference(avatar_component_clause, [], [f3806])).
fof(f3806, plain, (spl42_512 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_512])])).
fof(f27598, plain, (~ (e20 = op2(e23, e23)) | (~ spl42_339 | ~ spl42_871 | ~ spl42_883 | spl42_1193)), inference(forward_demodulation, [], [f27597, f5818])).
fof(f5818, plain, ((e20 = h2(e12)) | ~ spl42_883), inference(avatar_component_clause, [], [f5817])).
fof(f5817, plain, (spl42_883 <=> (e20 = h2(e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_883])])).
fof(f27597, plain, (~ (op2(e23, e23) = h2(e12)) | (~ spl42_339 | ~ spl42_871 | spl42_1193)), inference(forward_demodulation, [], [f27596, f2993])).
fof(f2993, plain, ((op1(e10, e10) = e12) | ~ spl42_339), inference(avatar_component_clause, [], [f2991])).
fof(f2991, plain, (spl42_339 <=> (op1(e10, e10) = e12)), introduced(avatar_definition, [new_symbols(naming, [spl42_339])])).
fof(f27596, plain, (~ (op2(e23, e23) = h2(op1(e10, e10))) | (~ spl42_871 | spl42_1193)), inference(forward_demodulation, [], [f7884, f5757])).
fof(f7884, plain, (~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10))) | spl42_1193), inference(avatar_component_clause, [], [f7882])).
fof(f7882, plain, (spl42_1193 <=> (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1193])])).
fof(f27583, plain, (~ spl42_585 | ~ spl42_487), inference(avatar_split_clause, [], [f27582, f3702, f4108])).
fof(f4108, plain, (spl42_585 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_585])])).
fof(f3702, plain, (spl42_487 <=> (e23 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_487])])).
fof(f27582, plain, (~ (e23 = op2(e22, e20)) | ~ spl42_487), inference(forward_demodulation, [], [f809, f3704])).
fof(f3704, plain, ((e23 = op2(e24, e20)) | ~ spl42_487), inference(avatar_component_clause, [], [f3702])).
fof(f809, plain, ~ (op2(e22, e20) = op2(e24, e20)), inference(cnf_transformation, [], [f8])).
fof(f27579, plain, (~ spl42_584 | ~ spl42_542), inference(avatar_split_clause, [], [f27578, f3930, f4104])).
fof(f4104, plain, (spl42_584 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_584])])).
fof(f27578, plain, (~ (e22 = op2(e22, e20)) | ~ spl42_542), inference(forward_demodulation, [], [f991, f3932])).
fof(f991, plain, ~ (op2(e22, e20) = op2(e22, e26)), inference(cnf_transformation, [], [f8])).
fof(f27523, plain, (~ spl42_255 | spl42_1241), inference(avatar_contradiction_clause, [], [f27522])).
fof(f27522, plain, ($false | (~ spl42_255 | spl42_1241)), inference(trivial_inequality_removal, [], [f27521])).
fof(f27521, plain, (~ (h2(e12) = h2(e12)) | (~ spl42_255 | spl42_1241)), inference(forward_demodulation, [], [f8076, f2645])).
fof(f2645, plain, ((e12 = op1(e11, e15)) | ~ spl42_255), inference(avatar_component_clause, [], [f2643])).
fof(f2643, plain, (spl42_255 <=> (e12 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_255])])).
fof(f8076, plain, (~ (h2(e12) = h2(op1(e11, e15))) | spl42_1241), inference(avatar_component_clause, [], [f8074])).
fof(f8074, plain, (spl42_1241 <=> (h2(e12) = h2(op1(e11, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1241])])).
fof(f27478, plain, (~ spl42_152 | ~ spl42_579 | ~ spl42_860 | ~ spl42_875 | spl42_1238), inference(avatar_contradiction_clause, [], [f27477])).
fof(f27477, plain, ($false | (~ spl42_152 | ~ spl42_579 | ~ spl42_860 | ~ spl42_875 | spl42_1238)), inference(subsumption_resolution, [], [f27476, f4085])).
fof(f4085, plain, ((e24 = op2(e22, e21)) | ~ spl42_579), inference(avatar_component_clause, [], [f4083])).
fof(f4083, plain, (spl42_579 <=> (e24 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_579])])).
fof(f27476, plain, (~ (e24 = op2(e22, e21)) | (~ spl42_152 | ~ spl42_860 | ~ spl42_875 | spl42_1238)), inference(forward_demodulation, [], [f27475, f5702])).
fof(f5702, plain, ((e24 = h2(e14)) | ~ spl42_860), inference(avatar_component_clause, [], [f5701])).
fof(f27475, plain, (~ (op2(e22, e21) = h2(e14)) | (~ spl42_152 | ~ spl42_875 | spl42_1238)), inference(forward_demodulation, [], [f27474, f2218])).
fof(f2218, plain, ((e14 = op1(e13, e16)) | ~ spl42_152), inference(avatar_component_clause, [], [f2216])).
fof(f2216, plain, (spl42_152 <=> (e14 = op1(e13, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_152])])).
fof(f27474, plain, (~ (op2(e22, e21) = h2(op1(e13, e16))) | (~ spl42_875 | spl42_1238)), inference(forward_demodulation, [], [f8064, f5777])).
fof(f5777, plain, ((e22 = h2(e13)) | ~ spl42_875), inference(avatar_component_clause, [], [f5776])).
fof(f5776, plain, (spl42_875 <=> (e22 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_875])])).
fof(f8064, plain, (~ (h2(op1(e13, e16)) = op2(h2(e13), e21)) | spl42_1238), inference(avatar_component_clause, [], [f8062])).
fof(f8062, plain, (spl42_1238 <=> (h2(op1(e13, e16)) = op2(h2(e13), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_1238])])).
fof(f27428, plain, (~ spl42_100 | ~ spl42_482 | ~ spl42_856 | ~ spl42_860 | spl42_1237), inference(avatar_contradiction_clause, [], [f27427])).
fof(f27427, plain, ($false | (~ spl42_100 | ~ spl42_482 | ~ spl42_856 | ~ spl42_860 | spl42_1237)), inference(subsumption_resolution, [], [f27426, f3683])).
fof(f3683, plain, ((e25 = op2(e24, e21)) | ~ spl42_482), inference(avatar_component_clause, [], [f3681])).
fof(f3681, plain, (spl42_482 <=> (e25 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_482])])).
fof(f27426, plain, (~ (e25 = op2(e24, e21)) | (~ spl42_100 | ~ spl42_856 | ~ spl42_860 | spl42_1237)), inference(forward_demodulation, [], [f27425, f5682])).
fof(f5682, plain, ((e25 = h2(e11)) | ~ spl42_856), inference(avatar_component_clause, [], [f5681])).
fof(f5681, plain, (spl42_856 <=> (e25 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_856])])).
fof(f27425, plain, (~ (op2(e24, e21) = h2(e11)) | (~ spl42_100 | ~ spl42_860 | spl42_1237)), inference(forward_demodulation, [], [f27424, f2003])).
fof(f2003, plain, ((e11 = op1(e14, e16)) | ~ spl42_100), inference(avatar_component_clause, [], [f2001])).
fof(f2001, plain, (spl42_100 <=> (e11 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_100])])).
fof(f27424, plain, (~ (op2(e24, e21) = h2(op1(e14, e16))) | (~ spl42_860 | spl42_1237)), inference(forward_demodulation, [], [f8060, f5702])).
fof(f8060, plain, (~ (h2(op1(e14, e16)) = op2(h2(e14), e21)) | spl42_1237), inference(avatar_component_clause, [], [f8058])).
fof(f8058, plain, (spl42_1237 <=> (h2(op1(e14, e16)) = op2(h2(e14), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_1237])])).
fof(f27372, plain, (~ spl42_50 | ~ spl42_382 | ~ spl42_871 | ~ spl42_1227 | spl42_1235), inference(avatar_contradiction_clause, [], [f27371])).
fof(f27371, plain, ($false | (~ spl42_50 | ~ spl42_382 | ~ spl42_871 | ~ spl42_1227 | spl42_1235)), inference(subsumption_resolution, [], [f27370, f3269])).
fof(f3269, plain, ((e23 = op2(e26, e21)) | ~ spl42_382), inference(avatar_component_clause, [], [f3267])).
fof(f3267, plain, (spl42_382 <=> (e23 = op2(e26, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_382])])).
fof(f27370, plain, (~ (e23 = op2(e26, e21)) | (~ spl42_50 | ~ spl42_871 | ~ spl42_1227 | spl42_1235)), inference(forward_demodulation, [], [f27369, f5757])).
fof(f27369, plain, (~ (op2(e26, e21) = h2(e10)) | (~ spl42_50 | ~ spl42_1227 | spl42_1235)), inference(forward_demodulation, [], [f27368, f1796])).
fof(f1796, plain, ((e10 = op1(e15, e16)) | ~ spl42_50), inference(avatar_component_clause, [], [f1794])).
fof(f1794, plain, (spl42_50 <=> (e10 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_50])])).
fof(f27368, plain, (~ (op2(e26, e21) = h2(op1(e15, e16))) | (~ spl42_1227 | spl42_1235)), inference(forward_demodulation, [], [f8052, f8019])).
fof(f8052, plain, (~ (h2(op1(e15, e16)) = op2(h2(e15), e21)) | spl42_1235), inference(avatar_component_clause, [], [f8050])).
fof(f8050, plain, (spl42_1235 <=> (h2(op1(e15, e16)) = op2(h2(e15), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_1235])])).
fof(f27321, plain, (~ spl42_49 | ~ spl42_611 | ~ spl42_871 | spl42_1234), inference(avatar_contradiction_clause, [], [f27320])).
fof(f27320, plain, ($false | (~ spl42_49 | ~ spl42_611 | ~ spl42_871 | spl42_1234)), inference(subsumption_resolution, [], [f27319, f4218])).
fof(f27319, plain, (~ (e21 = op2(e21, e23)) | (~ spl42_49 | ~ spl42_871 | spl42_1234)), inference(forward_demodulation, [], [f27318, f1205])).
fof(f1205, plain, (e21 = h2(e16)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(e21, e21) = h2(e15)) & (h2(e14) = op2(e21, op2(e21, op2(op2(e21, e21), op2(e21, e21))))) & (h2(e13) = op2(e21, op2(op2(e21, e21), op2(e21, e21)))) & (h2(e12) = op2(op2(op2(e21, e21), op2(e21, e21)), op2(e21, e21))) & (h2(e11) = op2(op2(e21, e21), op2(e21, e21))) & (h2(e10) = op2(e21, op2(e21, e21))) & (e21 = h2(e16))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax17)).
fof(f27318, plain, (~ (op2(e21, e23) = h2(e16)) | (~ spl42_49 | ~ spl42_871 | spl42_1234)), inference(forward_demodulation, [], [f27317, f1791])).
fof(f1791, plain, ((e16 = op1(e16, e10)) | ~ spl42_49), inference(avatar_component_clause, [], [f1789])).
fof(f1789, plain, (spl42_49 <=> (e16 = op1(e16, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_49])])).
fof(f27317, plain, (~ (op2(e21, e23) = h2(op1(e16, e10))) | (~ spl42_871 | spl42_1234)), inference(forward_demodulation, [], [f8048, f5757])).
fof(f8048, plain, (~ (h2(op1(e16, e10)) = op2(e21, h2(e10))) | spl42_1234), inference(avatar_component_clause, [], [f8046])).
fof(f8046, plain, (spl42_1234 <=> (h2(op1(e16, e10)) = op2(e21, h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1234])])).
fof(f27277, plain, (~ spl42_31 | ~ spl42_631 | ~ spl42_883 | spl42_1232), inference(avatar_contradiction_clause, [], [f27276])).
fof(f27276, plain, ($false | (~ spl42_31 | ~ spl42_631 | ~ spl42_883 | spl42_1232)), inference(subsumption_resolution, [], [f27275, f4301])).
fof(f4301, plain, ((e20 = op2(e21, e20)) | ~ spl42_631), inference(avatar_component_clause, [], [f4299])).
fof(f4299, plain, (spl42_631 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_631])])).
fof(f27275, plain, (~ (e20 = op2(e21, e20)) | (~ spl42_31 | ~ spl42_883 | spl42_1232)), inference(forward_demodulation, [], [f27274, f5818])).
fof(f27274, plain, (~ (op2(e21, e20) = h2(e12)) | (~ spl42_31 | ~ spl42_883 | spl42_1232)), inference(forward_demodulation, [], [f27273, f1717])).
fof(f1717, plain, ((e12 = op1(e16, e12)) | ~ spl42_31), inference(avatar_component_clause, [], [f1715])).
fof(f1715, plain, (spl42_31 <=> (e12 = op1(e16, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_31])])).
fof(f27273, plain, (~ (op2(e21, e20) = h2(op1(e16, e12))) | (~ spl42_883 | spl42_1232)), inference(forward_demodulation, [], [f8040, f5818])).
fof(f8040, plain, (~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | spl42_1232), inference(avatar_component_clause, [], [f8038])).
fof(f8038, plain, (spl42_1232 <=> (h2(op1(e16, e12)) = op2(e21, h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1232])])).
fof(f27226, plain, (~ spl42_26 | spl42_1231), inference(avatar_contradiction_clause, [], [f27225])).
fof(f27225, plain, ($false | (~ spl42_26 | spl42_1231)), inference(trivial_inequality_removal, [], [f27224])).
fof(f27224, plain, (~ (h2(e14) = h2(e14)) | (~ spl42_26 | spl42_1231)), inference(forward_demodulation, [], [f8036, f1696])).
fof(f1696, plain, ((e14 = op1(e16, e13)) | ~ spl42_26), inference(avatar_component_clause, [], [f1694])).
fof(f1694, plain, (spl42_26 <=> (e14 = op1(e16, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_26])])).
fof(f8036, plain, (~ (h2(e14) = h2(op1(e16, e13))) | spl42_1231), inference(avatar_component_clause, [], [f8034])).
fof(f8034, plain, (spl42_1231 <=> (h2(e14) = h2(op1(e16, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1231])])).
fof(f27176, plain, (~ spl42_16 | ~ spl42_608 | ~ spl42_856 | ~ spl42_860 | spl42_1230), inference(avatar_contradiction_clause, [], [f27175])).
fof(f27175, plain, ($false | (~ spl42_16 | ~ spl42_608 | ~ spl42_856 | ~ spl42_860 | spl42_1230)), inference(subsumption_resolution, [], [f27174, f4205])).
fof(f4205, plain, ((e25 = op2(e21, e24)) | ~ spl42_608), inference(avatar_component_clause, [], [f4203])).
fof(f4203, plain, (spl42_608 <=> (e25 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_608])])).
fof(f27174, plain, (~ (e25 = op2(e21, e24)) | (~ spl42_16 | ~ spl42_856 | ~ spl42_860 | spl42_1230)), inference(forward_demodulation, [], [f27173, f5682])).
fof(f27173, plain, (~ (op2(e21, e24) = h2(e11)) | (~ spl42_16 | ~ spl42_860 | spl42_1230)), inference(forward_demodulation, [], [f27172, f1655])).
fof(f1655, plain, ((e11 = op1(e16, e14)) | ~ spl42_16), inference(avatar_component_clause, [], [f1653])).
fof(f1653, plain, (spl42_16 <=> (e11 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_16])])).
fof(f27172, plain, (~ (op2(e21, e24) = h2(op1(e16, e14))) | (~ spl42_860 | spl42_1230)), inference(forward_demodulation, [], [f8032, f5702])).
fof(f8032, plain, (~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | spl42_1230), inference(avatar_component_clause, [], [f8030])).
fof(f8030, plain, (spl42_1230 <=> (h2(op1(e16, e14)) = op2(e21, h2(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1230])])).
fof(f27123, plain, (~ spl42_6 | spl42_1228), inference(avatar_contradiction_clause, [], [f27122])).
fof(f27122, plain, ($false | (~ spl42_6 | spl42_1228)), inference(trivial_inequality_removal, [], [f27121])).
fof(f27121, plain, (~ (h2(e15) = h2(e15)) | (~ spl42_6 | spl42_1228)), inference(forward_demodulation, [], [f8024, f1613])).
fof(f1613, plain, ((e15 = op1(e16, e16)) | ~ spl42_6), inference(avatar_component_clause, [], [f1611])).
fof(f1611, plain, (spl42_6 <=> (e15 = op1(e16, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_6])])).
fof(f8024, plain, (~ (h2(e15) = h2(op1(e16, e16))) | spl42_1228), inference(avatar_component_clause, [], [f8022])).
fof(f8022, plain, (spl42_1228 <=> (h2(e15) = h2(op1(e16, e16)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1228])])).
fof(f27072, plain, (~ spl42_74 | ~ spl42_374 | ~ spl42_875 | spl42_1225 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f27071])).
fof(f27071, plain, ($false | (~ spl42_74 | ~ spl42_374 | ~ spl42_875 | spl42_1225 | ~ spl42_1227)), inference(subsumption_resolution, [], [f27070, f3236])).
fof(f3236, plain, ((e22 = op2(e26, e22)) | ~ spl42_374), inference(avatar_component_clause, [], [f3234])).
fof(f3234, plain, (spl42_374 <=> (e22 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_374])])).
fof(f27070, plain, (~ (e22 = op2(e26, e22)) | (~ spl42_74 | ~ spl42_875 | spl42_1225 | ~ spl42_1227)), inference(forward_demodulation, [], [f27069, f5777])).
fof(f27069, plain, (~ (op2(e26, e22) = h2(e13)) | (~ spl42_74 | ~ spl42_875 | spl42_1225 | ~ spl42_1227)), inference(forward_demodulation, [], [f27068, f1895])).
fof(f1895, plain, ((e13 = op1(e15, e13)) | ~ spl42_74), inference(avatar_component_clause, [], [f1893])).
fof(f1893, plain, (spl42_74 <=> (e13 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_74])])).
fof(f27068, plain, (~ (op2(e26, e22) = h2(op1(e15, e13))) | (~ spl42_875 | spl42_1225 | ~ spl42_1227)), inference(forward_demodulation, [], [f27067, f8019])).
fof(f27067, plain, (~ (h2(op1(e15, e13)) = op2(h2(e15), e22)) | (~ spl42_875 | spl42_1225)), inference(forward_demodulation, [], [f8012, f5777])).
fof(f8012, plain, (~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | spl42_1225), inference(avatar_component_clause, [], [f8010])).
fof(f8010, plain, (spl42_1225 <=> (h2(op1(e15, e13)) = op2(h2(e15), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1225])])).
fof(f27017, plain, (~ spl42_83 | ~ spl42_392 | ~ spl42_883 | spl42_1224 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f27016])).
fof(f27016, plain, ($false | (~ spl42_83 | ~ spl42_392 | ~ spl42_883 | spl42_1224 | ~ spl42_1227)), inference(subsumption_resolution, [], [f27015, f3310])).
fof(f27015, plain, (~ (e26 = op2(e26, e20)) | (~ spl42_83 | ~ spl42_883 | spl42_1224 | ~ spl42_1227)), inference(forward_demodulation, [], [f27014, f8019])).
fof(f27014, plain, (~ (op2(e26, e20) = h2(e15)) | (~ spl42_83 | ~ spl42_883 | spl42_1224 | ~ spl42_1227)), inference(forward_demodulation, [], [f27013, f1932])).
fof(f1932, plain, ((e15 = op1(e15, e12)) | ~ spl42_83), inference(avatar_component_clause, [], [f1930])).
fof(f1930, plain, (spl42_83 <=> (e15 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_83])])).
fof(f27013, plain, (~ (op2(e26, e20) = h2(op1(e15, e12))) | (~ spl42_883 | spl42_1224 | ~ spl42_1227)), inference(forward_demodulation, [], [f27012, f8019])).
fof(f27012, plain, (~ (h2(op1(e15, e12)) = op2(h2(e15), e20)) | (~ spl42_883 | spl42_1224)), inference(forward_demodulation, [], [f8008, f5818])).
fof(f8008, plain, (~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | spl42_1224), inference(avatar_component_clause, [], [f8006])).
fof(f8006, plain, (spl42_1224 <=> (h2(op1(e15, e12)) = op2(h2(e15), h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1224])])).
fof(f26957, plain, (~ spl42_87 | ~ spl42_351 | ~ spl42_856 | ~ spl42_883 | spl42_1223 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f26956])).
fof(f26956, plain, ($false | (~ spl42_87 | ~ spl42_351 | ~ spl42_856 | ~ spl42_883 | spl42_1223 | ~ spl42_1227)), inference(subsumption_resolution, [], [f26955, f3141])).
fof(f3141, plain, ((e20 = op2(e26, e25)) | ~ spl42_351), inference(avatar_component_clause, [], [f3139])).
fof(f3139, plain, (spl42_351 <=> (e20 = op2(e26, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_351])])).
fof(f26955, plain, (~ (e20 = op2(e26, e25)) | (~ spl42_87 | ~ spl42_856 | ~ spl42_883 | spl42_1223 | ~ spl42_1227)), inference(forward_demodulation, [], [f26954, f5818])).
fof(f26954, plain, (~ (op2(e26, e25) = h2(e12)) | (~ spl42_87 | ~ spl42_856 | spl42_1223 | ~ spl42_1227)), inference(forward_demodulation, [], [f26953, f1949])).
fof(f1949, plain, ((e12 = op1(e15, e11)) | ~ spl42_87), inference(avatar_component_clause, [], [f1947])).
fof(f1947, plain, (spl42_87 <=> (e12 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_87])])).
fof(f26953, plain, (~ (op2(e26, e25) = h2(op1(e15, e11))) | (~ spl42_856 | spl42_1223 | ~ spl42_1227)), inference(forward_demodulation, [], [f26952, f8019])).
fof(f26952, plain, (~ (h2(op1(e15, e11)) = op2(h2(e15), e25)) | (~ spl42_856 | spl42_1223)), inference(forward_demodulation, [], [f8004, f5682])).
fof(f8004, plain, (~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | spl42_1223), inference(avatar_component_clause, [], [f8002])).
fof(f8002, plain, (spl42_1223 <=> (h2(op1(e15, e11)) = op2(h2(e15), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1223])])).
fof(f26893, plain, (~ spl42_96 | ~ spl42_369 | ~ spl42_860 | ~ spl42_871 | spl42_1222 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f26892])).
fof(f26892, plain, ($false | (~ spl42_96 | ~ spl42_369 | ~ spl42_860 | ~ spl42_871 | spl42_1222 | ~ spl42_1227)), inference(subsumption_resolution, [], [f26891, f3215])).
fof(f26891, plain, (~ (e24 = op2(e26, e23)) | (~ spl42_96 | ~ spl42_860 | ~ spl42_871 | spl42_1222 | ~ spl42_1227)), inference(forward_demodulation, [], [f26890, f5702])).
fof(f26890, plain, (~ (op2(e26, e23) = h2(e14)) | (~ spl42_96 | ~ spl42_871 | spl42_1222 | ~ spl42_1227)), inference(forward_demodulation, [], [f26889, f1986])).
fof(f1986, plain, ((e14 = op1(e15, e10)) | ~ spl42_96), inference(avatar_component_clause, [], [f1984])).
fof(f1984, plain, (spl42_96 <=> (e14 = op1(e15, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_96])])).
fof(f26889, plain, (~ (op2(e26, e23) = h2(op1(e15, e10))) | (~ spl42_871 | spl42_1222 | ~ spl42_1227)), inference(forward_demodulation, [], [f26888, f8019])).
fof(f26888, plain, (~ (h2(op1(e15, e10)) = op2(h2(e15), e23)) | (~ spl42_871 | spl42_1222)), inference(forward_demodulation, [], [f8000, f5757])).
fof(f8000, plain, (~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | spl42_1222), inference(avatar_component_clause, [], [f7998])).
fof(f7998, plain, (spl42_1222 <=> (h2(op1(e15, e10)) = op2(h2(e15), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1222])])).
fof(f26832, plain, (~ spl42_122 | ~ spl42_470 | ~ spl42_860 | ~ spl42_875 | ~ spl42_883 | spl42_1219), inference(avatar_contradiction_clause, [], [f26831])).
fof(f26831, plain, ($false | (~ spl42_122 | ~ spl42_470 | ~ spl42_860 | ~ spl42_875 | ~ spl42_883 | spl42_1219)), inference(subsumption_resolution, [], [f26830, f3634])).
fof(f3634, plain, ((e20 = op2(e24, e22)) | ~ spl42_470), inference(avatar_component_clause, [], [f3632])).
fof(f3632, plain, (spl42_470 <=> (e20 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_470])])).
fof(f26830, plain, (~ (e20 = op2(e24, e22)) | (~ spl42_122 | ~ spl42_860 | ~ spl42_875 | ~ spl42_883 | spl42_1219)), inference(forward_demodulation, [], [f26829, f5818])).
fof(f26829, plain, (~ (op2(e24, e22) = h2(e12)) | (~ spl42_122 | ~ spl42_860 | ~ spl42_875 | spl42_1219)), inference(forward_demodulation, [], [f26828, f2094])).
fof(f2094, plain, ((e12 = op1(e14, e13)) | ~ spl42_122), inference(avatar_component_clause, [], [f2092])).
fof(f2092, plain, (spl42_122 <=> (e12 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_122])])).
fof(f26828, plain, (~ (op2(e24, e22) = h2(op1(e14, e13))) | (~ spl42_860 | ~ spl42_875 | spl42_1219)), inference(forward_demodulation, [], [f26827, f5702])).
fof(f26827, plain, (~ (h2(op1(e14, e13)) = op2(h2(e14), e22)) | (~ spl42_875 | spl42_1219)), inference(forward_demodulation, [], [f7988, f5777])).
fof(f7988, plain, (~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | spl42_1219), inference(avatar_component_clause, [], [f7986])).
fof(f7986, plain, (spl42_1219 <=> (h2(op1(e14, e13)) = op2(h2(e14), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1219])])).
fof(f26774, plain, (~ spl42_127 | ~ spl42_487 | ~ spl42_860 | ~ spl42_871 | ~ spl42_883 | spl42_1218), inference(avatar_contradiction_clause, [], [f26773])).
fof(f26773, plain, ($false | (~ spl42_127 | ~ spl42_487 | ~ spl42_860 | ~ spl42_871 | ~ spl42_883 | spl42_1218)), inference(subsumption_resolution, [], [f26772, f3704])).
fof(f26772, plain, (~ (e23 = op2(e24, e20)) | (~ spl42_127 | ~ spl42_860 | ~ spl42_871 | ~ spl42_883 | spl42_1218)), inference(forward_demodulation, [], [f26771, f5757])).
fof(f26771, plain, (~ (op2(e24, e20) = h2(e10)) | (~ spl42_127 | ~ spl42_860 | ~ spl42_883 | spl42_1218)), inference(forward_demodulation, [], [f26770, f2115])).
fof(f2115, plain, ((e10 = op1(e14, e12)) | ~ spl42_127), inference(avatar_component_clause, [], [f2113])).
fof(f2113, plain, (spl42_127 <=> (e10 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_127])])).
fof(f26770, plain, (~ (op2(e24, e20) = h2(op1(e14, e12))) | (~ spl42_860 | ~ spl42_883 | spl42_1218)), inference(forward_demodulation, [], [f26769, f5702])).
fof(f26769, plain, (~ (h2(op1(e14, e12)) = op2(h2(e14), e20)) | (~ spl42_883 | spl42_1218)), inference(forward_demodulation, [], [f7984, f5818])).
fof(f7984, plain, (~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | spl42_1218), inference(avatar_component_clause, [], [f7982])).
fof(f7982, plain, (spl42_1218 <=> (h2(op1(e14, e12)) = op2(h2(e14), h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1218])])).
fof(f26717, plain, (~ spl42_144 | ~ spl42_465 | ~ spl42_860 | ~ spl42_871 | ~ spl42_875 | spl42_1216), inference(avatar_contradiction_clause, [], [f26716])).
fof(f26716, plain, ($false | (~ spl42_144 | ~ spl42_465 | ~ spl42_860 | ~ spl42_871 | ~ spl42_875 | spl42_1216)), inference(subsumption_resolution, [], [f26715, f3613])).
fof(f26715, plain, (~ (e22 = op2(e24, e23)) | (~ spl42_144 | ~ spl42_860 | ~ spl42_871 | ~ spl42_875 | spl42_1216)), inference(forward_demodulation, [], [f26714, f5777])).
fof(f26714, plain, (~ (op2(e24, e23) = h2(e13)) | (~ spl42_144 | ~ spl42_860 | ~ spl42_871 | spl42_1216)), inference(forward_demodulation, [], [f26713, f2185])).
fof(f2185, plain, ((e13 = op1(e14, e10)) | ~ spl42_144), inference(avatar_component_clause, [], [f2183])).
fof(f2183, plain, (spl42_144 <=> (e13 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_144])])).
fof(f26713, plain, (~ (op2(e24, e23) = h2(op1(e14, e10))) | (~ spl42_860 | ~ spl42_871 | spl42_1216)), inference(forward_demodulation, [], [f26712, f5702])).
fof(f26712, plain, (~ (h2(op1(e14, e10)) = op2(h2(e14), e23)) | (~ spl42_871 | spl42_1216)), inference(forward_demodulation, [], [f7976, f5757])).
fof(f7976, plain, (~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | spl42_1216), inference(avatar_component_clause, [], [f7974])).
fof(f7974, plain, (spl42_1216 <=> (h2(op1(e14, e10)) = op2(h2(e14), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1216])])).
fof(f26657, plain, (~ spl42_58 | spl42_1236), inference(avatar_contradiction_clause, [], [f26656])).
fof(f26656, plain, ($false | (~ spl42_58 | spl42_1236)), inference(trivial_inequality_removal, [], [f26655])).
fof(f26655, plain, (~ (h2(e11) = h2(e11)) | (~ spl42_58 | spl42_1236)), inference(forward_demodulation, [], [f8056, f1829])).
fof(f1829, plain, ((e11 = op1(e15, e15)) | ~ spl42_58), inference(avatar_component_clause, [], [f1827])).
fof(f1827, plain, (spl42_58 <=> (e11 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_58])])).
fof(f8056, plain, (~ (h2(e11) = h2(op1(e15, e15))) | spl42_1236), inference(avatar_component_clause, [], [f8054])).
fof(f8054, plain, (spl42_1236 <=> (h2(e11) = h2(op1(e15, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1236])])).
fof(f26597, plain, (~ spl42_8 | spl42_1229), inference(avatar_contradiction_clause, [], [f26596])).
fof(f26596, plain, ($false | (~ spl42_8 | spl42_1229)), inference(trivial_inequality_removal, [], [f26595])).
fof(f26595, plain, (~ (h2(e10) = h2(e10)) | (~ spl42_8 | spl42_1229)), inference(forward_demodulation, [], [f8028, f1622])).
fof(f1622, plain, ((e10 = op1(e16, e15)) | ~ spl42_8), inference(avatar_component_clause, [], [f1620])).
fof(f1620, plain, (spl42_8 <=> (e10 = op1(e16, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_8])])).
fof(f8028, plain, (~ (h2(e10) = h2(op1(e16, e15))) | spl42_1229), inference(avatar_component_clause, [], [f8026])).
fof(f8026, plain, (spl42_1229 <=> (h2(e10) = h2(op1(e16, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1229])])).
fof(f26542, plain, (~ spl42_112 | ~ spl42_443 | ~ spl42_860 | spl42_1221 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f26541])).
fof(f26541, plain, ($false | (~ spl42_112 | ~ spl42_443 | ~ spl42_860 | spl42_1221 | ~ spl42_1227)), inference(subsumption_resolution, [], [f26540, f3522])).
fof(f26540, plain, (~ (e21 = op2(e24, e26)) | (~ spl42_112 | ~ spl42_860 | spl42_1221 | ~ spl42_1227)), inference(forward_demodulation, [], [f26539, f1205])).
fof(f26539, plain, (~ (op2(e24, e26) = h2(e16)) | (~ spl42_112 | ~ spl42_860 | spl42_1221 | ~ spl42_1227)), inference(forward_demodulation, [], [f26538, f2052])).
fof(f2052, plain, ((e16 = op1(e14, e15)) | ~ spl42_112), inference(avatar_component_clause, [], [f2050])).
fof(f2050, plain, (spl42_112 <=> (e16 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_112])])).
fof(f26538, plain, (~ (op2(e24, e26) = h2(op1(e14, e15))) | (~ spl42_860 | spl42_1221 | ~ spl42_1227)), inference(forward_demodulation, [], [f26537, f5702])).
fof(f26537, plain, (~ (h2(op1(e14, e15)) = op2(h2(e14), e26)) | (spl42_1221 | ~ spl42_1227)), inference(forward_demodulation, [], [f7996, f8019])).
fof(f7996, plain, (~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | spl42_1221), inference(avatar_component_clause, [], [f7994])).
fof(f7994, plain, (spl42_1221 <=> (h2(op1(e14, e15)) = op2(h2(e14), h2(e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1221])])).
fof(f26492, plain, (~ spl42_158 | ~ spl42_542 | ~ spl42_875 | spl42_1215 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f26491])).
fof(f26491, plain, ($false | (~ spl42_158 | ~ spl42_542 | ~ spl42_875 | spl42_1215 | ~ spl42_1227)), inference(subsumption_resolution, [], [f26490, f3932])).
fof(f26490, plain, (~ (e22 = op2(e22, e26)) | (~ spl42_158 | ~ spl42_875 | spl42_1215 | ~ spl42_1227)), inference(forward_demodulation, [], [f26489, f5777])).
fof(f26489, plain, (~ (op2(e22, e26) = h2(e13)) | (~ spl42_158 | ~ spl42_875 | spl42_1215 | ~ spl42_1227)), inference(forward_demodulation, [], [f26488, f2243])).
fof(f2243, plain, ((e13 = op1(e13, e15)) | ~ spl42_158), inference(avatar_component_clause, [], [f2241])).
fof(f2241, plain, (spl42_158 <=> (e13 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_158])])).
fof(f26488, plain, (~ (op2(e22, e26) = h2(op1(e13, e15))) | (~ spl42_875 | spl42_1215 | ~ spl42_1227)), inference(forward_demodulation, [], [f26487, f5777])).
fof(f26487, plain, (~ (h2(op1(e13, e15)) = op2(h2(e13), e26)) | (spl42_1215 | ~ spl42_1227)), inference(forward_demodulation, [], [f7972, f8019])).
fof(f7972, plain, (~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | spl42_1215), inference(avatar_component_clause, [], [f7970])).
fof(f7970, plain, (spl42_1215 <=> (h2(op1(e13, e15)) = op2(h2(e13), h2(e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1215])])).
fof(f26437, plain, (~ spl42_169 | ~ spl42_571 | ~ spl42_871 | ~ spl42_875 | spl42_1213), inference(avatar_contradiction_clause, [], [f26436])).
fof(f26436, plain, ($false | (~ spl42_169 | ~ spl42_571 | ~ spl42_871 | ~ spl42_875 | spl42_1213)), inference(subsumption_resolution, [], [f26435, f4052])).
fof(f4052, plain, ((e23 = op2(e22, e22)) | ~ spl42_571), inference(avatar_component_clause, [], [f4050])).
fof(f4050, plain, (spl42_571 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_571])])).
fof(f26435, plain, (~ (e23 = op2(e22, e22)) | (~ spl42_169 | ~ spl42_871 | ~ spl42_875 | spl42_1213)), inference(forward_demodulation, [], [f26434, f5757])).
fof(f26434, plain, (~ (op2(e22, e22) = h2(e10)) | (~ spl42_169 | ~ spl42_875 | spl42_1213)), inference(forward_demodulation, [], [f26433, f2289])).
fof(f2289, plain, ((e10 = op1(e13, e13)) | ~ spl42_169), inference(avatar_component_clause, [], [f2287])).
fof(f2287, plain, (spl42_169 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_169])])).
fof(f26433, plain, (~ (op2(e22, e22) = h2(op1(e13, e13))) | (~ spl42_875 | spl42_1213)), inference(forward_demodulation, [], [f7964, f5777])).
fof(f7964, plain, (~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | spl42_1213), inference(avatar_component_clause, [], [f7962])).
fof(f7962, plain, (spl42_1213 <=> (h2(op1(e13, e13)) = op2(h2(e13), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1213])])).
fof(f26379, plain, (~ spl42_182 | ~ spl42_583 | ~ spl42_875 | ~ spl42_883 | spl42_1212), inference(avatar_contradiction_clause, [], [f26378])).
fof(f26378, plain, ($false | (~ spl42_182 | ~ spl42_583 | ~ spl42_875 | ~ spl42_883 | spl42_1212)), inference(subsumption_resolution, [], [f26377, f4102])).
fof(f4102, plain, ((e21 = op2(e22, e20)) | ~ spl42_583), inference(avatar_component_clause, [], [f4100])).
fof(f4100, plain, (spl42_583 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_583])])).
fof(f26377, plain, (~ (e21 = op2(e22, e20)) | (~ spl42_182 | ~ spl42_875 | ~ spl42_883 | spl42_1212)), inference(forward_demodulation, [], [f26376, f1205])).
fof(f26376, plain, (~ (op2(e22, e20) = h2(e16)) | (~ spl42_182 | ~ spl42_875 | ~ spl42_883 | spl42_1212)), inference(forward_demodulation, [], [f26375, f2342])).
fof(f2342, plain, ((e16 = op1(e13, e12)) | ~ spl42_182), inference(avatar_component_clause, [], [f2340])).
fof(f2340, plain, (spl42_182 <=> (e16 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_182])])).
fof(f26375, plain, (~ (op2(e22, e20) = h2(op1(e13, e12))) | (~ spl42_875 | ~ spl42_883 | spl42_1212)), inference(forward_demodulation, [], [f26374, f5777])).
fof(f26374, plain, (~ (h2(op1(e13, e12)) = op2(h2(e13), e20)) | (~ spl42_883 | spl42_1212)), inference(forward_demodulation, [], [f7960, f5818])).
fof(f7960, plain, (~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | spl42_1212), inference(avatar_component_clause, [], [f7958])).
fof(f7958, plain, (spl42_1212 <=> (h2(op1(e13, e12)) = op2(h2(e13), h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1212])])).
fof(f26315, plain, (~ spl42_184 | ~ spl42_552 | ~ spl42_856 | ~ spl42_875 | spl42_1211), inference(avatar_contradiction_clause, [], [f26314])).
fof(f26314, plain, ($false | (~ spl42_184 | ~ spl42_552 | ~ spl42_856 | ~ spl42_875 | spl42_1211)), inference(subsumption_resolution, [], [f26313, f3973])).
fof(f26313, plain, (~ (e25 = op2(e22, e25)) | (~ spl42_184 | ~ spl42_856 | ~ spl42_875 | spl42_1211)), inference(forward_demodulation, [], [f26312, f5682])).
fof(f26312, plain, (~ (op2(e22, e25) = h2(e11)) | (~ spl42_184 | ~ spl42_856 | ~ spl42_875 | spl42_1211)), inference(forward_demodulation, [], [f26311, f2351])).
fof(f2351, plain, ((e11 = op1(e13, e11)) | ~ spl42_184), inference(avatar_component_clause, [], [f2349])).
fof(f2349, plain, (spl42_184 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_184])])).
fof(f26311, plain, (~ (op2(e22, e25) = h2(op1(e13, e11))) | (~ spl42_856 | ~ spl42_875 | spl42_1211)), inference(forward_demodulation, [], [f26310, f5777])).
fof(f26310, plain, (~ (h2(op1(e13, e11)) = op2(h2(e13), e25)) | (~ spl42_856 | spl42_1211)), inference(forward_demodulation, [], [f7956, f5682])).
fof(f7956, plain, (~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | spl42_1211), inference(avatar_component_clause, [], [f7954])).
fof(f7954, plain, (spl42_1211 <=> (h2(op1(e13, e11)) = op2(h2(e13), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1211])])).
fof(f26249, plain, (~ spl42_195 | ~ spl42_567 | ~ spl42_871 | ~ spl42_875 | spl42_1210 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f26248])).
fof(f26248, plain, ($false | (~ spl42_195 | ~ spl42_567 | ~ spl42_871 | ~ spl42_875 | spl42_1210 | ~ spl42_1227)), inference(subsumption_resolution, [], [f26247, f4035])).
fof(f4035, plain, ((e26 = op2(e22, e23)) | ~ spl42_567), inference(avatar_component_clause, [], [f4033])).
fof(f26247, plain, (~ (e26 = op2(e22, e23)) | (~ spl42_195 | ~ spl42_871 | ~ spl42_875 | spl42_1210 | ~ spl42_1227)), inference(forward_demodulation, [], [f26246, f8019])).
fof(f26246, plain, (~ (op2(e22, e23) = h2(e15)) | (~ spl42_195 | ~ spl42_871 | ~ spl42_875 | spl42_1210)), inference(forward_demodulation, [], [f26245, f2396])).
fof(f2396, plain, ((e15 = op1(e13, e10)) | ~ spl42_195), inference(avatar_component_clause, [], [f2394])).
fof(f2394, plain, (spl42_195 <=> (e15 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_195])])).
fof(f26245, plain, (~ (op2(e22, e23) = h2(op1(e13, e10))) | (~ spl42_871 | ~ spl42_875 | spl42_1210)), inference(forward_demodulation, [], [f26244, f5777])).
fof(f26244, plain, (~ (h2(op1(e13, e10)) = op2(h2(e13), e23)) | (~ spl42_871 | spl42_1210)), inference(forward_demodulation, [], [f7952, f5757])).
fof(f7952, plain, (~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | spl42_1210), inference(avatar_component_clause, [], [f7950])).
fof(f7950, plain, (spl42_1210 <=> (h2(op1(e13, e10)) = op2(h2(e13), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1210])])).
fof(f26184, plain, (~ spl42_211 | ~ spl42_655 | ~ spl42_860 | ~ spl42_871 | ~ spl42_883 | spl42_1208), inference(avatar_contradiction_clause, [], [f26183])).
fof(f26183, plain, ($false | (~ spl42_211 | ~ spl42_655 | ~ spl42_860 | ~ spl42_871 | ~ spl42_883 | spl42_1208)), inference(subsumption_resolution, [], [f26182, f4400])).
fof(f4400, plain, ((e23 = op2(e20, e24)) | ~ spl42_655), inference(avatar_component_clause, [], [f4398])).
fof(f4398, plain, (spl42_655 <=> (e23 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_655])])).
fof(f26182, plain, (~ (e23 = op2(e20, e24)) | (~ spl42_211 | ~ spl42_860 | ~ spl42_871 | ~ spl42_883 | spl42_1208)), inference(forward_demodulation, [], [f26181, f5757])).
fof(f26181, plain, (~ (op2(e20, e24) = h2(e10)) | (~ spl42_211 | ~ spl42_860 | ~ spl42_883 | spl42_1208)), inference(forward_demodulation, [], [f26180, f2463])).
fof(f2463, plain, ((e10 = op1(e12, e14)) | ~ spl42_211), inference(avatar_component_clause, [], [f2461])).
fof(f2461, plain, (spl42_211 <=> (e10 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_211])])).
fof(f26180, plain, (~ (op2(e20, e24) = h2(op1(e12, e14))) | (~ spl42_860 | ~ spl42_883 | spl42_1208)), inference(forward_demodulation, [], [f26179, f5818])).
fof(f26179, plain, (~ (h2(op1(e12, e14)) = op2(h2(e12), e24)) | (~ spl42_860 | spl42_1208)), inference(forward_demodulation, [], [f7944, f5702])).
fof(f7944, plain, (~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | spl42_1208), inference(avatar_component_clause, [], [f7942])).
fof(f7942, plain, (spl42_1208 <=> (h2(op1(e12, e14)) = op2(h2(e12), h2(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1208])])).
fof(f26119, plain, (~ spl42_224 | ~ spl42_667 | ~ spl42_875 | ~ spl42_883 | spl42_1207), inference(avatar_contradiction_clause, [], [f26118])).
fof(f26118, plain, ($false | (~ spl42_224 | ~ spl42_667 | ~ spl42_875 | ~ spl42_883 | spl42_1207)), inference(subsumption_resolution, [], [f26117, f4450])).
fof(f26117, plain, (~ (e21 = op2(e20, e22)) | (~ spl42_224 | ~ spl42_875 | ~ spl42_883 | spl42_1207)), inference(forward_demodulation, [], [f26116, f1205])).
fof(f26116, plain, (~ (op2(e20, e22) = h2(e16)) | (~ spl42_224 | ~ spl42_875 | ~ spl42_883 | spl42_1207)), inference(forward_demodulation, [], [f26115, f2516])).
fof(f2516, plain, ((e16 = op1(e12, e13)) | ~ spl42_224), inference(avatar_component_clause, [], [f2514])).
fof(f2514, plain, (spl42_224 <=> (e16 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_224])])).
fof(f26115, plain, (~ (op2(e20, e22) = h2(op1(e12, e13))) | (~ spl42_875 | ~ spl42_883 | spl42_1207)), inference(forward_demodulation, [], [f26114, f5818])).
fof(f26114, plain, (~ (h2(op1(e12, e13)) = op2(h2(e12), e22)) | (~ spl42_875 | spl42_1207)), inference(forward_demodulation, [], [f7940, f5777])).
fof(f7940, plain, (~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | spl42_1207), inference(avatar_component_clause, [], [f7938])).
fof(f7938, plain, (spl42_1207 <=> (h2(op1(e12, e13)) = op2(h2(e12), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1207])])).
fof(f26053, plain, (~ spl42_228 | ~ spl42_682 | ~ spl42_875 | ~ spl42_883 | spl42_1206), inference(avatar_contradiction_clause, [], [f26052])).
fof(f26052, plain, ($false | (~ spl42_228 | ~ spl42_682 | ~ spl42_875 | ~ spl42_883 | spl42_1206)), inference(subsumption_resolution, [], [f26051, f4512])).
fof(f4512, plain, ((op2(e20, e20) = e22) | ~ spl42_682), inference(avatar_component_clause, [], [f4510])).
fof(f4510, plain, (spl42_682 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl42_682])])).
fof(f26051, plain, (~ (op2(e20, e20) = e22) | (~ spl42_228 | ~ spl42_875 | ~ spl42_883 | spl42_1206)), inference(forward_demodulation, [], [f26050, f5777])).
fof(f26050, plain, (~ (op2(e20, e20) = h2(e13)) | (~ spl42_228 | ~ spl42_883 | spl42_1206)), inference(forward_demodulation, [], [f26049, f2533])).
fof(f2533, plain, ((e13 = op1(e12, e12)) | ~ spl42_228), inference(avatar_component_clause, [], [f2531])).
fof(f2531, plain, (spl42_228 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_228])])).
fof(f26049, plain, (~ (op2(e20, e20) = h2(op1(e12, e12))) | (~ spl42_883 | spl42_1206)), inference(forward_demodulation, [], [f7936, f5818])).
fof(f7936, plain, (~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | spl42_1206), inference(avatar_component_clause, [], [f7934])).
fof(f7934, plain, (spl42_1206 <=> (h2(op1(e12, e12)) = op2(h2(e12), h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1206])])).
fof(f25989, plain, (~ spl42_236 | ~ spl42_649 | ~ spl42_856 | ~ spl42_860 | ~ spl42_883 | spl42_1205), inference(avatar_contradiction_clause, [], [f25988])).
fof(f25988, plain, ($false | (~ spl42_236 | ~ spl42_649 | ~ spl42_856 | ~ spl42_860 | ~ spl42_883 | spl42_1205)), inference(subsumption_resolution, [], [f25987, f4375])).
fof(f4375, plain, ((e24 = op2(e20, e25)) | ~ spl42_649), inference(avatar_component_clause, [], [f4373])).
fof(f4373, plain, (spl42_649 <=> (e24 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_649])])).
fof(f25987, plain, (~ (e24 = op2(e20, e25)) | (~ spl42_236 | ~ spl42_856 | ~ spl42_860 | ~ spl42_883 | spl42_1205)), inference(forward_demodulation, [], [f25986, f5702])).
fof(f25986, plain, (~ (op2(e20, e25) = h2(e14)) | (~ spl42_236 | ~ spl42_856 | ~ spl42_883 | spl42_1205)), inference(forward_demodulation, [], [f25985, f2566])).
fof(f2566, plain, ((e14 = op1(e12, e11)) | ~ spl42_236), inference(avatar_component_clause, [], [f2564])).
fof(f2564, plain, (spl42_236 <=> (e14 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_236])])).
fof(f25985, plain, (~ (op2(e20, e25) = h2(op1(e12, e11))) | (~ spl42_856 | ~ spl42_883 | spl42_1205)), inference(forward_demodulation, [], [f25984, f5818])).
fof(f25984, plain, (~ (h2(op1(e12, e11)) = op2(h2(e12), e25)) | (~ spl42_856 | spl42_1205)), inference(forward_demodulation, [], [f7932, f5682])).
fof(f7932, plain, (~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | spl42_1205), inference(avatar_component_clause, [], [f7930])).
fof(f7930, plain, (spl42_1205 <=> (h2(op1(e12, e11)) = op2(h2(e12), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1205])])).
fof(f25909, plain, (~ spl42_240 | ~ spl42_664 | ~ spl42_856 | ~ spl42_871 | ~ spl42_883 | spl42_1204), inference(avatar_contradiction_clause, [], [f25908])).
fof(f25908, plain, ($false | (~ spl42_240 | ~ spl42_664 | ~ spl42_856 | ~ spl42_871 | ~ spl42_883 | spl42_1204)), inference(subsumption_resolution, [], [f25907, f4437])).
fof(f4437, plain, ((e25 = op2(e20, e23)) | ~ spl42_664), inference(avatar_component_clause, [], [f4435])).
fof(f4435, plain, (spl42_664 <=> (e25 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_664])])).
fof(f25907, plain, (~ (e25 = op2(e20, e23)) | (~ spl42_240 | ~ spl42_856 | ~ spl42_871 | ~ spl42_883 | spl42_1204)), inference(forward_demodulation, [], [f25906, f5682])).
fof(f25906, plain, (~ (op2(e20, e23) = h2(e11)) | (~ spl42_240 | ~ spl42_871 | ~ spl42_883 | spl42_1204)), inference(forward_demodulation, [], [f25905, f2583])).
fof(f2583, plain, ((e11 = op1(e12, e10)) | ~ spl42_240), inference(avatar_component_clause, [], [f2581])).
fof(f2581, plain, (spl42_240 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_240])])).
fof(f25905, plain, (~ (op2(e20, e23) = h2(op1(e12, e10))) | (~ spl42_871 | ~ spl42_883 | spl42_1204)), inference(forward_demodulation, [], [f25904, f5818])).
fof(f25904, plain, (~ (h2(op1(e12, e10)) = op2(h2(e12), e23)) | (~ spl42_871 | spl42_1204)), inference(forward_demodulation, [], [f7928, f5757])).
fof(f7928, plain, (~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | spl42_1204), inference(avatar_component_clause, [], [f7926])).
fof(f7926, plain, (spl42_1204 <=> (h2(op1(e12, e10)) = op2(h2(e12), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1204])])).
fof(f25886, plain, (~ spl42_628 | ~ spl42_630), inference(avatar_contradiction_clause, [], [f25885])).
fof(f25885, plain, ($false | (~ spl42_628 | ~ spl42_630)), inference(subsumption_resolution, [], [f25884, f1131])).
fof(f25884, plain, ((e24 = e26) | (~ spl42_628 | ~ spl42_630)), inference(backward_demodulation, [], [f4296, f4288])).
fof(f4288, plain, ((e24 = op2(e21, e21)) | ~ spl42_628), inference(avatar_component_clause, [], [f4286])).
fof(f4286, plain, (spl42_628 <=> (e24 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_628])])).
fof(f4296, plain, ((e26 = op2(e21, e21)) | ~ spl42_630), inference(avatar_component_clause, [], [f4294])).
fof(f4294, plain, (spl42_630 <=> (e26 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_630])])).
fof(f25807, plain, (~ spl42_622 | ~ spl42_426), inference(avatar_split_clause, [], [f21796, f3449, f4261])).
fof(f4261, plain, (spl42_622 <=> (e25 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_622])])).
fof(f21796, plain, (~ (e25 = op2(e21, e22)) | ~ spl42_426), inference(forward_demodulation, [], [f848, f3451])).
fof(f848, plain, ~ (op2(e21, e22) = op2(e25, e22)), inference(cnf_transformation, [], [f8])).
fof(f25757, plain, (~ spl42_265 | ~ spl42_413 | ~ spl42_856 | ~ spl42_860 | spl42_1203 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f25756])).
fof(f25756, plain, ($false | (~ spl42_265 | ~ spl42_413 | ~ spl42_856 | ~ spl42_860 | spl42_1203 | ~ spl42_1227)), inference(subsumption_resolution, [], [f25755, f3397])).
fof(f3397, plain, ((e26 = op2(e25, e24)) | ~ spl42_413), inference(avatar_component_clause, [], [f3395])).
fof(f25755, plain, (~ (e26 = op2(e25, e24)) | (~ spl42_265 | ~ spl42_856 | ~ spl42_860 | spl42_1203 | ~ spl42_1227)), inference(forward_demodulation, [], [f25754, f8019])).
fof(f25754, plain, (~ (op2(e25, e24) = h2(e15)) | (~ spl42_265 | ~ spl42_856 | ~ spl42_860 | spl42_1203)), inference(forward_demodulation, [], [f25753, f2686])).
fof(f2686, plain, ((e15 = op1(e11, e14)) | ~ spl42_265), inference(avatar_component_clause, [], [f2684])).
fof(f2684, plain, (spl42_265 <=> (e15 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_265])])).
fof(f25753, plain, (~ (op2(e25, e24) = h2(op1(e11, e14))) | (~ spl42_856 | ~ spl42_860 | spl42_1203)), inference(forward_demodulation, [], [f25752, f5682])).
fof(f25752, plain, (~ (h2(op1(e11, e14)) = op2(h2(e11), e24)) | (~ spl42_860 | spl42_1203)), inference(forward_demodulation, [], [f7924, f5702])).
fof(f7924, plain, (~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | spl42_1203), inference(avatar_component_clause, [], [f7922])).
fof(f7922, plain, (spl42_1203 <=> (h2(op1(e11, e14)) = op2(h2(e11), h2(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1203])])).
fof(f25695, plain, (~ spl42_268 | ~ spl42_426 | ~ spl42_856 | ~ spl42_875 | spl42_1202), inference(avatar_contradiction_clause, [], [f25694])).
fof(f25694, plain, ($false | (~ spl42_268 | ~ spl42_426 | ~ spl42_856 | ~ spl42_875 | spl42_1202)), inference(subsumption_resolution, [], [f25693, f3451])).
fof(f25693, plain, (~ (e25 = op2(e25, e22)) | (~ spl42_268 | ~ spl42_856 | ~ spl42_875 | spl42_1202)), inference(forward_demodulation, [], [f25692, f5682])).
fof(f25692, plain, (~ (op2(e25, e22) = h2(e11)) | (~ spl42_268 | ~ spl42_856 | ~ spl42_875 | spl42_1202)), inference(forward_demodulation, [], [f25691, f2699])).
fof(f2699, plain, ((e11 = op1(e11, e13)) | ~ spl42_268), inference(avatar_component_clause, [], [f2697])).
fof(f2697, plain, (spl42_268 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_268])])).
fof(f25691, plain, (~ (op2(e25, e22) = h2(op1(e11, e13))) | (~ spl42_856 | ~ spl42_875 | spl42_1202)), inference(forward_demodulation, [], [f25690, f5682])).
fof(f25690, plain, (~ (h2(op1(e11, e13)) = op2(h2(e11), e22)) | (~ spl42_875 | spl42_1202)), inference(forward_demodulation, [], [f7920, f5777])).
fof(f7920, plain, (~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | spl42_1202), inference(avatar_component_clause, [], [f7918])).
fof(f7918, plain, (spl42_1202 <=> (h2(op1(e11, e13)) = op2(h2(e11), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1202])])).
fof(f25636, plain, (~ spl42_287 | ~ spl42_401 | ~ spl42_856 | spl42_1200), inference(avatar_contradiction_clause, [], [f25635])).
fof(f25635, plain, ($false | (~ spl42_287 | ~ spl42_401 | ~ spl42_856 | spl42_1200)), inference(subsumption_resolution, [], [f25634, f3348])).
fof(f25634, plain, (~ (e21 = op2(e25, e25)) | (~ spl42_287 | ~ spl42_856 | spl42_1200)), inference(forward_demodulation, [], [f25633, f1205])).
fof(f25633, plain, (~ (op2(e25, e25) = h2(e16)) | (~ spl42_287 | ~ spl42_856 | spl42_1200)), inference(forward_demodulation, [], [f25632, f2777])).
fof(f2777, plain, ((e16 = op1(e11, e11)) | ~ spl42_287), inference(avatar_component_clause, [], [f2775])).
fof(f2775, plain, (spl42_287 <=> (e16 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_287])])).
fof(f25632, plain, (~ (op2(e25, e25) = h2(op1(e11, e11))) | (~ spl42_856 | spl42_1200)), inference(forward_demodulation, [], [f7912, f5682])).
fof(f7912, plain, (~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | spl42_1200), inference(avatar_component_clause, [], [f7910])).
fof(f7910, plain, (spl42_1200 <=> (h2(op1(e11, e11)) = op2(h2(e11), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1200])])).
fof(f25577, plain, (~ spl42_288 | ~ spl42_417 | ~ spl42_856 | ~ spl42_871 | spl42_1199), inference(avatar_contradiction_clause, [], [f25576])).
fof(f25576, plain, ($false | (~ spl42_288 | ~ spl42_417 | ~ spl42_856 | ~ spl42_871 | spl42_1199)), inference(subsumption_resolution, [], [f25575, f3414])).
fof(f3414, plain, ((e23 = op2(e25, e23)) | ~ spl42_417), inference(avatar_component_clause, [], [f3412])).
fof(f3412, plain, (spl42_417 <=> (e23 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_417])])).
fof(f25575, plain, (~ (e23 = op2(e25, e23)) | (~ spl42_288 | ~ spl42_856 | ~ spl42_871 | spl42_1199)), inference(forward_demodulation, [], [f25574, f5757])).
fof(f25574, plain, (~ (op2(e25, e23) = h2(e10)) | (~ spl42_288 | ~ spl42_856 | ~ spl42_871 | spl42_1199)), inference(forward_demodulation, [], [f25573, f2782])).
fof(f2782, plain, ((e10 = op1(e11, e10)) | ~ spl42_288), inference(avatar_component_clause, [], [f2780])).
fof(f2780, plain, (spl42_288 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_288])])).
fof(f25573, plain, (~ (op2(e25, e23) = h2(op1(e11, e10))) | (~ spl42_856 | ~ spl42_871 | spl42_1199)), inference(forward_demodulation, [], [f25572, f5682])).
fof(f25572, plain, (~ (h2(op1(e11, e10)) = op2(h2(e11), e23)) | (~ spl42_871 | spl42_1199)), inference(forward_demodulation, [], [f7908, f5757])).
fof(f7908, plain, (~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | spl42_1199), inference(avatar_component_clause, [], [f7906])).
fof(f7906, plain, (spl42_1199 <=> (h2(op1(e11, e10)) = op2(h2(e11), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1199])])).
fof(f25453, plain, (~ spl42_312 | ~ spl42_507 | ~ spl42_860 | ~ spl42_871 | ~ spl42_875 | spl42_1197), inference(avatar_contradiction_clause, [], [f25452])).
fof(f25452, plain, ($false | (~ spl42_312 | ~ spl42_507 | ~ spl42_860 | ~ spl42_871 | ~ spl42_875 | spl42_1197)), inference(subsumption_resolution, [], [f25451, f3787])).
fof(f3787, plain, ((e22 = op2(e23, e24)) | ~ spl42_507), inference(avatar_component_clause, [], [f3785])).
fof(f3785, plain, (spl42_507 <=> (e22 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_507])])).
fof(f25451, plain, (~ (e22 = op2(e23, e24)) | (~ spl42_312 | ~ spl42_860 | ~ spl42_871 | ~ spl42_875 | spl42_1197)), inference(forward_demodulation, [], [f25450, f5777])).
fof(f25450, plain, (~ (op2(e23, e24) = h2(e13)) | (~ spl42_312 | ~ spl42_860 | ~ spl42_871 | spl42_1197)), inference(forward_demodulation, [], [f25449, f2881])).
fof(f2881, plain, ((e13 = op1(e10, e14)) | ~ spl42_312), inference(avatar_component_clause, [], [f2879])).
fof(f2879, plain, (spl42_312 <=> (e13 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_312])])).
fof(f25449, plain, (~ (op2(e23, e24) = h2(op1(e10, e14))) | (~ spl42_860 | ~ spl42_871 | spl42_1197)), inference(forward_demodulation, [], [f25448, f5757])).
fof(f25448, plain, (~ (h2(op1(e10, e14)) = op2(h2(e10), e24)) | (~ spl42_860 | spl42_1197)), inference(forward_demodulation, [], [f7900, f5702])).
fof(f7900, plain, (~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | spl42_1197), inference(avatar_component_clause, [], [f7898])).
fof(f7898, plain, (spl42_1197 <=> (h2(op1(e10, e14)) = op2(h2(e10), h2(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1197])])).
fof(f25391, plain, (~ spl42_321 | ~ spl42_525 | ~ spl42_871 | ~ spl42_875 | spl42_1196 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f25390])).
fof(f25390, plain, ($false | (~ spl42_321 | ~ spl42_525 | ~ spl42_871 | ~ spl42_875 | spl42_1196 | ~ spl42_1227)), inference(subsumption_resolution, [], [f25389, f3861])).
fof(f3861, plain, ((e26 = op2(e23, e22)) | ~ spl42_525), inference(avatar_component_clause, [], [f3859])).
fof(f3859, plain, (spl42_525 <=> (e26 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_525])])).
fof(f25389, plain, (~ (e26 = op2(e23, e22)) | (~ spl42_321 | ~ spl42_871 | ~ spl42_875 | spl42_1196 | ~ spl42_1227)), inference(forward_demodulation, [], [f25388, f8019])).
fof(f25388, plain, (~ (op2(e23, e22) = h2(e15)) | (~ spl42_321 | ~ spl42_871 | ~ spl42_875 | spl42_1196)), inference(forward_demodulation, [], [f25387, f2918])).
fof(f2918, plain, ((e15 = op1(e10, e13)) | ~ spl42_321), inference(avatar_component_clause, [], [f2916])).
fof(f2916, plain, (spl42_321 <=> (e15 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_321])])).
fof(f25387, plain, (~ (op2(e23, e22) = h2(op1(e10, e13))) | (~ spl42_871 | ~ spl42_875 | spl42_1196)), inference(forward_demodulation, [], [f25386, f5757])).
fof(f25386, plain, (~ (h2(op1(e10, e13)) = op2(h2(e10), e22)) | (~ spl42_875 | spl42_1196)), inference(forward_demodulation, [], [f7896, f5777])).
fof(f7896, plain, (~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | spl42_1196), inference(avatar_component_clause, [], [f7894])).
fof(f7894, plain, (spl42_1196 <=> (h2(op1(e10, e13)) = op2(h2(e10), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1196])])).
fof(f25325, plain, (~ spl42_301 | ~ spl42_527 | ~ spl42_871 | spl42_1242), inference(avatar_contradiction_clause, [], [f25324])).
fof(f25324, plain, ($false | (~ spl42_301 | ~ spl42_527 | ~ spl42_871 | spl42_1242)), inference(subsumption_resolution, [], [f25323, f3870])).
fof(f3870, plain, ((e21 = op2(e23, e21)) | ~ spl42_527), inference(avatar_component_clause, [], [f3868])).
fof(f3868, plain, (spl42_527 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_527])])).
fof(f25323, plain, (~ (e21 = op2(e23, e21)) | (~ spl42_301 | ~ spl42_871 | spl42_1242)), inference(forward_demodulation, [], [f25322, f1205])).
fof(f25322, plain, (~ (op2(e23, e21) = h2(e16)) | (~ spl42_301 | ~ spl42_871 | spl42_1242)), inference(forward_demodulation, [], [f25321, f2835])).
fof(f2835, plain, ((e16 = op1(e10, e16)) | ~ spl42_301), inference(avatar_component_clause, [], [f2833])).
fof(f2833, plain, (spl42_301 <=> (e16 = op1(e10, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_301])])).
fof(f25321, plain, (~ (op2(e23, e21) = h2(op1(e10, e16))) | (~ spl42_871 | spl42_1242)), inference(forward_demodulation, [], [f8080, f5757])).
fof(f8080, plain, (~ (h2(op1(e10, e16)) = op2(h2(e10), e21)) | spl42_1242), inference(avatar_component_clause, [], [f8078])).
fof(f8078, plain, (spl42_1242 <=> (h2(op1(e10, e16)) = op2(h2(e10), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_1242])])).
fof(f25261, plain, (~ spl42_39 | spl42_1233), inference(avatar_contradiction_clause, [], [f25260])).
fof(f25260, plain, ($false | (~ spl42_39 | spl42_1233)), inference(trivial_inequality_removal, [], [f25259])).
fof(f25259, plain, (~ (h2(e13) = h2(e13)) | (~ spl42_39 | spl42_1233)), inference(forward_demodulation, [], [f8044, f1750])).
fof(f1750, plain, ((e13 = op1(e16, e11)) | ~ spl42_39), inference(avatar_component_clause, [], [f1748])).
fof(f1748, plain, (spl42_39 <=> (e13 = op1(e16, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_39])])).
fof(f8044, plain, (~ (h2(e13) = h2(op1(e16, e11))) | spl42_1233), inference(avatar_component_clause, [], [f8042])).
fof(f8042, plain, (spl42_1233 <=> (h2(e13) = h2(op1(e16, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1233])])).
fof(f25200, plain, (~ spl42_70 | ~ spl42_359 | ~ spl42_860 | spl42_1226 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f25199])).
fof(f25199, plain, ($false | (~ spl42_70 | ~ spl42_359 | ~ spl42_860 | spl42_1226 | ~ spl42_1227)), inference(subsumption_resolution, [], [f25198, f3174])).
fof(f3174, plain, ((e21 = op2(e26, e24)) | ~ spl42_359), inference(avatar_component_clause, [], [f3172])).
fof(f3172, plain, (spl42_359 <=> (e21 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_359])])).
fof(f25198, plain, (~ (e21 = op2(e26, e24)) | (~ spl42_70 | ~ spl42_860 | spl42_1226 | ~ spl42_1227)), inference(forward_demodulation, [], [f25197, f1205])).
fof(f25197, plain, (~ (op2(e26, e24) = h2(e16)) | (~ spl42_70 | ~ spl42_860 | spl42_1226 | ~ spl42_1227)), inference(forward_demodulation, [], [f25196, f1878])).
fof(f1878, plain, ((e16 = op1(e15, e14)) | ~ spl42_70), inference(avatar_component_clause, [], [f1876])).
fof(f1876, plain, (spl42_70 <=> (e16 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_70])])).
fof(f25196, plain, (~ (op2(e26, e24) = h2(op1(e15, e14))) | (~ spl42_860 | spl42_1226 | ~ spl42_1227)), inference(forward_demodulation, [], [f25195, f8019])).
fof(f25195, plain, (~ (h2(op1(e15, e14)) = op2(h2(e15), e24)) | (~ spl42_860 | spl42_1226)), inference(forward_demodulation, [], [f8016, f5702])).
fof(f8016, plain, (~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | spl42_1226), inference(avatar_component_clause, [], [f8014])).
fof(f8014, plain, (spl42_1226 <=> (h2(op1(e15, e14)) = op2(h2(e15), h2(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1226])])).
fof(f25134, plain, (~ spl42_117 | ~ spl42_460 | ~ spl42_860 | spl42_1220), inference(avatar_contradiction_clause, [], [f25133])).
fof(f25133, plain, ($false | (~ spl42_117 | ~ spl42_460 | ~ spl42_860 | spl42_1220)), inference(subsumption_resolution, [], [f25132, f3592])).
fof(f3592, plain, ((e24 = op2(e24, e24)) | ~ spl42_460), inference(avatar_component_clause, [], [f3590])).
fof(f3590, plain, (spl42_460 <=> (e24 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_460])])).
fof(f25132, plain, (~ (e24 = op2(e24, e24)) | (~ spl42_117 | ~ spl42_860 | spl42_1220)), inference(forward_demodulation, [], [f25131, f5702])).
fof(f25131, plain, (~ (op2(e24, e24) = h2(e14)) | (~ spl42_117 | ~ spl42_860 | spl42_1220)), inference(forward_demodulation, [], [f25130, f2073])).
fof(f2073, plain, ((e14 = op1(e14, e14)) | ~ spl42_117), inference(avatar_component_clause, [], [f2071])).
fof(f2071, plain, (spl42_117 <=> (e14 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_117])])).
fof(f25130, plain, (~ (op2(e24, e24) = h2(op1(e14, e14))) | (~ spl42_860 | spl42_1220)), inference(forward_demodulation, [], [f7992, f5702])).
fof(f7992, plain, (~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | spl42_1220), inference(avatar_component_clause, [], [f7990])).
fof(f7990, plain, (spl42_1220 <=> (h2(op1(e14, e14)) = op2(h2(e14), h2(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1220])])).
fof(f25072, plain, (~ spl42_164 | ~ spl42_554 | ~ spl42_860 | ~ spl42_875 | ~ spl42_883 | spl42_1214), inference(avatar_contradiction_clause, [], [f25071])).
fof(f25071, plain, ($false | (~ spl42_164 | ~ spl42_554 | ~ spl42_860 | ~ spl42_875 | ~ spl42_883 | spl42_1214)), inference(subsumption_resolution, [], [f25070, f3982])).
fof(f3982, plain, ((e20 = op2(e22, e24)) | ~ spl42_554), inference(avatar_component_clause, [], [f3980])).
fof(f3980, plain, (spl42_554 <=> (e20 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_554])])).
fof(f25070, plain, (~ (e20 = op2(e22, e24)) | (~ spl42_164 | ~ spl42_860 | ~ spl42_875 | ~ spl42_883 | spl42_1214)), inference(forward_demodulation, [], [f25069, f5818])).
fof(f25069, plain, (~ (op2(e22, e24) = h2(e12)) | (~ spl42_164 | ~ spl42_860 | ~ spl42_875 | spl42_1214)), inference(forward_demodulation, [], [f25068, f2268])).
fof(f2268, plain, ((e12 = op1(e13, e14)) | ~ spl42_164), inference(avatar_component_clause, [], [f2266])).
fof(f2266, plain, (spl42_164 <=> (e12 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_164])])).
fof(f25068, plain, (~ (op2(e22, e24) = h2(op1(e13, e14))) | (~ spl42_860 | ~ spl42_875 | spl42_1214)), inference(forward_demodulation, [], [f25067, f5777])).
fof(f25067, plain, (~ (h2(op1(e13, e14)) = op2(h2(e13), e24)) | (~ spl42_860 | spl42_1214)), inference(forward_demodulation, [], [f7968, f5702])).
fof(f7968, plain, (~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | spl42_1214), inference(avatar_component_clause, [], [f7966])).
fof(f7966, plain, (spl42_1214 <=> (h2(op1(e13, e14)) = op2(h2(e13), h2(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1214])])).
fof(f25012, plain, (~ spl42_209 | ~ spl42_644 | ~ spl42_883 | spl42_1209 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f25011])).
fof(f25011, plain, ($false | (~ spl42_209 | ~ spl42_644 | ~ spl42_883 | spl42_1209 | ~ spl42_1227)), inference(subsumption_resolution, [], [f25010, f4354])).
fof(f4354, plain, ((e26 = op2(e20, e26)) | ~ spl42_644), inference(avatar_component_clause, [], [f4352])).
fof(f25010, plain, (~ (e26 = op2(e20, e26)) | (~ spl42_209 | ~ spl42_883 | spl42_1209 | ~ spl42_1227)), inference(forward_demodulation, [], [f25009, f8019])).
fof(f25009, plain, (~ (op2(e20, e26) = h2(e15)) | (~ spl42_209 | ~ spl42_883 | spl42_1209 | ~ spl42_1227)), inference(forward_demodulation, [], [f25008, f2454])).
fof(f2454, plain, ((e15 = op1(e12, e15)) | ~ spl42_209), inference(avatar_component_clause, [], [f2452])).
fof(f2452, plain, (spl42_209 <=> (e15 = op1(e12, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_209])])).
fof(f25008, plain, (~ (op2(e20, e26) = h2(op1(e12, e15))) | (~ spl42_883 | spl42_1209 | ~ spl42_1227)), inference(forward_demodulation, [], [f25007, f5818])).
fof(f25007, plain, (~ (h2(op1(e12, e15)) = op2(h2(e12), e26)) | (spl42_1209 | ~ spl42_1227)), inference(forward_demodulation, [], [f7948, f8019])).
fof(f7948, plain, (~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | spl42_1209), inference(avatar_component_clause, [], [f7946])).
fof(f7946, plain, (spl42_1209 <=> (h2(op1(e12, e15)) = op2(h2(e12), h2(e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1209])])).
fof(f24949, plain, (~ spl42_278 | ~ spl42_439 | ~ spl42_856 | ~ spl42_860 | ~ spl42_883 | spl42_1201), inference(avatar_contradiction_clause, [], [f24948])).
fof(f24948, plain, ($false | (~ spl42_278 | ~ spl42_439 | ~ spl42_856 | ~ spl42_860 | ~ spl42_883 | spl42_1201)), inference(subsumption_resolution, [], [f24947, f3505])).
fof(f3505, plain, ((e24 = op2(e25, e20)) | ~ spl42_439), inference(avatar_component_clause, [], [f3503])).
fof(f3503, plain, (spl42_439 <=> (e24 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_439])])).
fof(f24947, plain, (~ (e24 = op2(e25, e20)) | (~ spl42_278 | ~ spl42_856 | ~ spl42_860 | ~ spl42_883 | spl42_1201)), inference(forward_demodulation, [], [f24946, f5702])).
fof(f24946, plain, (~ (op2(e25, e20) = h2(e14)) | (~ spl42_278 | ~ spl42_856 | ~ spl42_883 | spl42_1201)), inference(forward_demodulation, [], [f24945, f2740])).
fof(f2740, plain, ((e14 = op1(e11, e12)) | ~ spl42_278), inference(avatar_component_clause, [], [f2738])).
fof(f2738, plain, (spl42_278 <=> (e14 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_278])])).
fof(f24945, plain, (~ (op2(e25, e20) = h2(op1(e11, e12))) | (~ spl42_856 | ~ spl42_883 | spl42_1201)), inference(forward_demodulation, [], [f24944, f5682])).
fof(f24944, plain, (~ (h2(op1(e11, e12)) = op2(h2(e11), e20)) | (~ spl42_883 | spl42_1201)), inference(forward_demodulation, [], [f7916, f5818])).
fof(f7916, plain, (~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | spl42_1201), inference(avatar_component_clause, [], [f7914])).
fof(f7914, plain, (spl42_1201 <=> (h2(op1(e11, e12)) = op2(h2(e11), h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1201])])).
fof(f24887, plain, (~ spl42_324 | ~ spl42_538 | ~ spl42_856 | ~ spl42_871 | ~ spl42_883 | spl42_1195), inference(avatar_contradiction_clause, [], [f24886])).
fof(f24886, plain, ($false | (~ spl42_324 | ~ spl42_538 | ~ spl42_856 | ~ spl42_871 | ~ spl42_883 | spl42_1195)), inference(subsumption_resolution, [], [f24885, f3915])).
fof(f3915, plain, ((e25 = op2(e23, e20)) | ~ spl42_538), inference(avatar_component_clause, [], [f3913])).
fof(f3913, plain, (spl42_538 <=> (e25 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_538])])).
fof(f24885, plain, (~ (e25 = op2(e23, e20)) | (~ spl42_324 | ~ spl42_856 | ~ spl42_871 | ~ spl42_883 | spl42_1195)), inference(forward_demodulation, [], [f24884, f5682])).
fof(f24884, plain, (~ (op2(e23, e20) = h2(e11)) | (~ spl42_324 | ~ spl42_871 | ~ spl42_883 | spl42_1195)), inference(forward_demodulation, [], [f24883, f2931])).
fof(f2931, plain, ((e11 = op1(e10, e12)) | ~ spl42_324), inference(avatar_component_clause, [], [f2929])).
fof(f2929, plain, (spl42_324 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_324])])).
fof(f24883, plain, (~ (op2(e23, e20) = h2(op1(e10, e12))) | (~ spl42_871 | ~ spl42_883 | spl42_1195)), inference(forward_demodulation, [], [f24882, f5757])).
fof(f24882, plain, (~ (h2(op1(e10, e12)) = op2(h2(e10), e20)) | (~ spl42_883 | spl42_1195)), inference(forward_demodulation, [], [f7892, f5818])).
fof(f7892, plain, (~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | spl42_1195), inference(avatar_component_clause, [], [f7890])).
fof(f7890, plain, (spl42_1195 <=> (h2(op1(e10, e12)) = op2(h2(e10), h2(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1195])])).
fof(f24814, plain, (~ spl42_249 | ~ spl42_430 | ~ spl42_856 | ~ spl42_875 | spl42_1240), inference(avatar_contradiction_clause, [], [f24813])).
fof(f24813, plain, ($false | (~ spl42_249 | ~ spl42_430 | ~ spl42_856 | ~ spl42_875 | spl42_1240)), inference(subsumption_resolution, [], [f24812, f3468])).
fof(f3468, plain, ((e22 = op2(e25, e21)) | ~ spl42_430), inference(avatar_component_clause, [], [f3466])).
fof(f3466, plain, (spl42_430 <=> (e22 = op2(e25, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_430])])).
fof(f24812, plain, (~ (e22 = op2(e25, e21)) | (~ spl42_249 | ~ spl42_856 | ~ spl42_875 | spl42_1240)), inference(forward_demodulation, [], [f24811, f5777])).
fof(f24811, plain, (~ (op2(e25, e21) = h2(e13)) | (~ spl42_249 | ~ spl42_856 | spl42_1240)), inference(forward_demodulation, [], [f24810, f2620])).
fof(f2620, plain, ((e13 = op1(e11, e16)) | ~ spl42_249), inference(avatar_component_clause, [], [f2618])).
fof(f2618, plain, (spl42_249 <=> (e13 = op1(e11, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_249])])).
fof(f24810, plain, (~ (op2(e25, e21) = h2(op1(e11, e16))) | (~ spl42_856 | spl42_1240)), inference(forward_demodulation, [], [f8072, f5682])).
fof(f8072, plain, (~ (h2(op1(e11, e16)) = op2(h2(e11), e21)) | spl42_1240), inference(avatar_component_clause, [], [f8070])).
fof(f8070, plain, (spl42_1240 <=> (h2(op1(e11, e16)) = op2(h2(e11), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_1240])])).
fof(f24725, plain, (~ spl42_139 | ~ spl42_455 | ~ spl42_856 | ~ spl42_860 | spl42_1217 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f24724])).
fof(f24724, plain, ($false | (~ spl42_139 | ~ spl42_455 | ~ spl42_856 | ~ spl42_860 | spl42_1217 | ~ spl42_1227)), inference(subsumption_resolution, [], [f24723, f3571])).
fof(f3571, plain, ((e26 = op2(e24, e25)) | ~ spl42_455), inference(avatar_component_clause, [], [f3569])).
fof(f3569, plain, (spl42_455 <=> (e26 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_455])])).
fof(f24723, plain, (~ (e26 = op2(e24, e25)) | (~ spl42_139 | ~ spl42_856 | ~ spl42_860 | spl42_1217 | ~ spl42_1227)), inference(forward_demodulation, [], [f24722, f8019])).
fof(f24722, plain, (~ (op2(e24, e25) = h2(e15)) | (~ spl42_139 | ~ spl42_856 | ~ spl42_860 | spl42_1217)), inference(forward_demodulation, [], [f24721, f2164])).
fof(f2164, plain, ((e15 = op1(e14, e11)) | ~ spl42_139), inference(avatar_component_clause, [], [f2162])).
fof(f2162, plain, (spl42_139 <=> (e15 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_139])])).
fof(f24721, plain, (~ (op2(e24, e25) = h2(op1(e14, e11))) | (~ spl42_856 | ~ spl42_860 | spl42_1217)), inference(forward_demodulation, [], [f24720, f5702])).
fof(f24720, plain, (~ (h2(op1(e14, e11)) = op2(h2(e14), e25)) | (~ spl42_856 | spl42_1217)), inference(forward_demodulation, [], [f7980, f5682])).
fof(f7980, plain, (~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | spl42_1217), inference(avatar_component_clause, [], [f7978])).
fof(f7978, plain, (spl42_1217 <=> (h2(op1(e14, e11)) = op2(h2(e14), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1217])])).
fof(f24696, plain, (~ spl42_184 | ~ spl42_185), inference(avatar_contradiction_clause, [], [f24695])).
fof(f24695, plain, ($false | (~ spl42_184 | ~ spl42_185)), inference(subsumption_resolution, [], [f24694, f1097])).
fof(f1097, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (~ (e15 = e16) & ~ (e14 = e16) & ~ (e14 = e15) & ~ (e13 = e16) & ~ (e13 = e15) & ~ (e13 = e14) & ~ (e12 = e16) & ~ (e12 = e15) & ~ (e12 = e14) & ~ (e12 = e13) & ~ (e11 = e16) & ~ (e11 = e15) & ~ (e11 = e14) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e16) & ~ (e10 = e15) & ~ (e10 = e14) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax9)).
fof(f24694, plain, ((e11 = e12) | (~ spl42_184 | ~ spl42_185)), inference(forward_demodulation, [], [f2355, f2351])).
fof(f2355, plain, ((e12 = op1(e13, e11)) | ~ spl42_185), inference(avatar_component_clause, [], [f2353])).
fof(f2353, plain, (spl42_185 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_185])])).
fof(f24670, plain, (spl42_191 | ~ spl42_330 | ~ spl42_340), inference(avatar_contradiction_clause, [], [f24669])).
fof(f24669, plain, ($false | (spl42_191 | ~ spl42_330 | ~ spl42_340)), inference(subsumption_resolution, [], [f24662, f2379])).
fof(f2379, plain, (~ (e11 = op1(e13, e10)) | spl42_191), inference(avatar_component_clause, [], [f2378])).
fof(f2378, plain, (spl42_191 <=> (e11 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_191])])).
fof(f24662, plain, ((e11 = op1(e13, e10)) | (~ spl42_330 | ~ spl42_340)), inference(backward_demodulation, [], [f24176, f2997])).
fof(f2997, plain, ((op1(e10, e10) = e13) | ~ spl42_340), inference(avatar_component_clause, [], [f2995])).
fof(f2995, plain, (spl42_340 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl42_340])])).
fof(f24176, plain, ((e11 = op1(op1(e10, e10), e10)) | ~ spl42_330), inference(backward_demodulation, [], [f265, f2956])).
fof(f2956, plain, ((e10 = op1(e10, e11)) | ~ spl42_330), inference(avatar_component_clause, [], [f2954])).
fof(f2954, plain, (spl42_330 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_330])])).
fof(f265, plain, (e11 = op1(op1(op1(e10, e11), e10), e10)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, ((e16 = op1(op1(op1(e16, e16), e16), e16)) & (e16 = op1(op1(op1(e15, e16), e15), e15)) & (e16 = op1(op1(op1(e14, e16), e14), e14)) & (e16 = op1(op1(op1(e13, e16), e13), e13)) & (e16 = op1(op1(op1(e12, e16), e12), e12)) & (e16 = op1(op1(op1(e11, e16), e11), e11)) & (e16 = op1(op1(op1(e10, e16), e10), e10)) & (e15 = op1(op1(op1(e16, e15), e16), e16)) & (e15 = op1(op1(op1(e15, e15), e15), e15)) & (e15 = op1(op1(op1(e14, e15), e14), e14)) & (e15 = op1(op1(op1(e13, e15), e13), e13)) & (e15 = op1(op1(op1(e12, e15), e12), e12)) & (e15 = op1(op1(op1(e11, e15), e11), e11)) & (e15 = op1(op1(op1(e10, e15), e10), e10)) & (e14 = op1(op1(op1(e16, e14), e16), e16)) & (e14 = op1(op1(op1(e15, e14), e15), e15)) & (e14 = op1(op1(op1(e14, e14), e14), e14)) & (e14 = op1(op1(op1(e13, e14), e13), e13)) & (e14 = op1(op1(op1(e12, e14), e12), e12)) & (e14 = op1(op1(op1(e11, e14), e11), e11)) & (e14 = op1(op1(op1(e10, e14), e10), e10)) & (e13 = op1(op1(op1(e16, e13), e16), e16)) & (e13 = op1(op1(op1(e15, e13), e15), e15)) & (e13 = op1(op1(op1(e14, e13), e14), e14)) & (e13 = op1(op1(op1(e13, e13), e13), e13)) & (e13 = op1(op1(op1(e12, e13), e12), e12)) & (e13 = op1(op1(op1(e11, e13), e11), e11)) & (e13 = op1(op1(op1(e10, e13), e10), e10)) & (e12 = op1(op1(op1(e16, e12), e16), e16)) & (e12 = op1(op1(op1(e15, e12), e15), e15)) & (e12 = op1(op1(op1(e14, e12), e14), e14)) & (e12 = op1(op1(op1(e13, e12), e13), e13)) & (e12 = op1(op1(op1(e12, e12), e12), e12)) & (e12 = op1(op1(op1(e11, e12), e11), e11)) & (e12 = op1(op1(op1(e10, e12), e10), e10)) & (e11 = op1(op1(op1(e16, e11), e16), e16)) & (e11 = op1(op1(op1(e15, e11), e15), e15)) & (e11 = op1(op1(op1(e14, e11), e14), e14)) & (e11 = op1(op1(op1(e13, e11), e13), e13)) & (e11 = op1(op1(op1(e12, e11), e12), e12)) & (e11 = op1(op1(op1(e11, e11), e11), e11)) & (e11 = op1(op1(op1(e10, e11), e10), e10)) & (e10 = op1(op1(op1(e16, e10), e16), e16)) & (e10 = op1(op1(op1(e15, e10), e15), e15)) & (e10 = op1(op1(op1(e14, e10), e14), e14)) & (e10 = op1(op1(op1(e13, e10), e13), e13)) & (e10 = op1(op1(op1(e12, e10), e12), e12)) & (e10 = op1(op1(op1(e11, e10), e11), e11)) & (e10 = op1(op1(op1(e10, e10), e10), e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax3)).
fof(f24627, plain, (spl42_470 | ~ spl42_583 | ~ spl42_860 | ~ spl42_875), inference(avatar_split_clause, [], [f24625, f5776, f5701, f4100, f3632])).
fof(f24625, plain, ((e20 = op2(e24, e22)) | (~ spl42_583 | ~ spl42_860 | ~ spl42_875)), inference(backward_demodulation, [], [f23901, f5702])).
fof(f23901, plain, ((e20 = op2(h2(e14), e22)) | (~ spl42_583 | ~ spl42_875)), inference(backward_demodulation, [], [f22083, f23895])).
fof(f23895, plain, ((op2(e21, e22) = h2(e14)) | ~ spl42_875), inference(backward_demodulation, [], [f4680, f5777])).
fof(f4680, plain, (h2(e14) = op2(e21, h2(e13))), inference(backward_demodulation, [], [f4678, f4679])).
fof(f4679, plain, (h2(e13) = op2(e21, op2(h2(e15), h2(e15)))), inference(forward_demodulation, [], [f1209, f1211])).
fof(f1211, plain, (op2(e21, e21) = h2(e15)), inference(cnf_transformation, [], [f17])).
fof(f1209, plain, (h2(e13) = op2(e21, op2(op2(e21, e21), op2(e21, e21)))), inference(cnf_transformation, [], [f17])).
fof(f4678, plain, (h2(e14) = op2(e21, op2(e21, op2(h2(e15), h2(e15))))), inference(forward_demodulation, [], [f1210, f1211])).
fof(f1210, plain, (h2(e14) = op2(e21, op2(e21, op2(op2(e21, e21), op2(e21, e21))))), inference(cnf_transformation, [], [f17])).
fof(f22083, plain, ((e20 = op2(op2(e21, e22), e22)) | ~ spl42_583), inference(backward_demodulation, [], [f456, f4102])).
fof(f456, plain, (e20 = op2(op2(op2(e22, e20), e22), e22)), inference(cnf_transformation, [], [f6])).
fof(f24606, plain, (~ spl42_199 | ~ spl42_673 | ~ spl42_883 | spl42_1239), inference(avatar_contradiction_clause, [], [f24605])).
fof(f24605, plain, ($false | (~ spl42_199 | ~ spl42_673 | ~ spl42_883 | spl42_1239)), inference(subsumption_resolution, [], [f24604, f4475])).
fof(f4475, plain, ((e20 = op2(e20, e21)) | ~ spl42_673), inference(avatar_component_clause, [], [f4473])).
fof(f4473, plain, (spl42_673 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_673])])).
fof(f24604, plain, (~ (e20 = op2(e20, e21)) | (~ spl42_199 | ~ spl42_883 | spl42_1239)), inference(forward_demodulation, [], [f24603, f5818])).
fof(f24603, plain, (~ (op2(e20, e21) = h2(e12)) | (~ spl42_199 | ~ spl42_883 | spl42_1239)), inference(forward_demodulation, [], [f24602, f2413])).
fof(f2413, plain, ((e12 = op1(e12, e16)) | ~ spl42_199), inference(avatar_component_clause, [], [f2411])).
fof(f2411, plain, (spl42_199 <=> (e12 = op1(e12, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_199])])).
fof(f24602, plain, (~ (op2(e20, e21) = h2(op1(e12, e16))) | (~ spl42_883 | spl42_1239)), inference(forward_demodulation, [], [f8068, f5818])).
fof(f8068, plain, (~ (h2(op1(e12, e16)) = op2(h2(e12), e21)) | spl42_1239), inference(avatar_component_clause, [], [f8066])).
fof(f8066, plain, (spl42_1239 <=> (h2(op1(e12, e16)) = op2(h2(e12), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_1239])])).
fof(f24542, plain, (spl42_571 | ~ spl42_374 | ~ spl42_567), inference(avatar_split_clause, [], [f24040, f4033, f3234, f4050])).
fof(f24040, plain, ((e23 = op2(e22, e22)) | (~ spl42_374 | ~ spl42_567)), inference(forward_demodulation, [], [f24036, f3236])).
fof(f24036, plain, ((e23 = op2(op2(e26, e22), e22)) | ~ spl42_567), inference(backward_demodulation, [], [f477, f4035])).
fof(f477, plain, (e23 = op2(op2(op2(e22, e23), e22), e22)), inference(cnf_transformation, [], [f6])).
fof(f24526, plain, (spl42_339 | ~ spl42_288 | ~ spl42_324), inference(avatar_split_clause, [], [f24254, f2929, f2780, f2991])).
fof(f24254, plain, ((op1(e10, e10) = e12) | (~ spl42_288 | ~ spl42_324)), inference(backward_demodulation, [], [f24198, f2782])).
fof(f24198, plain, ((e12 = op1(op1(e11, e10), e10)) | ~ spl42_324), inference(backward_demodulation, [], [f272, f2931])).
fof(f272, plain, (e12 = op1(op1(op1(e10, e12), e10), e10)), inference(cnf_transformation, [], [f3])).
fof(f24517, plain, (~ spl42_165 | ~ spl42_158), inference(avatar_split_clause, [], [f24516, f2241, f2270])).
fof(f2270, plain, (spl42_165 <=> (e13 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_165])])).
fof(f24516, plain, (~ (e13 = op1(e13, e14)) | ~ spl42_158), inference(forward_demodulation, [], [f731, f2243])).
fof(f731, plain, ~ (op1(e13, e14) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (op1(e16, e15) = op1(e16, e16)) & ~ (op1(e16, e14) = op1(e16, e16)) & ~ (op1(e16, e14) = op1(e16, e15)) & ~ (op1(e16, e13) = op1(e16, e16)) & ~ (op1(e16, e13) = op1(e16, e15)) & ~ (op1(e16, e13) = op1(e16, e14)) & ~ (op1(e16, e12) = op1(e16, e16)) & ~ (op1(e16, e12) = op1(e16, e15)) & ~ (op1(e16, e12) = op1(e16, e14)) & ~ (op1(e16, e12) = op1(e16, e13)) & ~ (op1(e16, e11) = op1(e16, e16)) & ~ (op1(e16, e11) = op1(e16, e15)) & ~ (op1(e16, e11) = op1(e16, e14)) & ~ (op1(e16, e11) = op1(e16, e13)) & ~ (op1(e16, e11) = op1(e16, e12)) & ~ (op1(e16, e10) = op1(e16, e16)) & ~ (op1(e16, e10) = op1(e16, e15)) & ~ (op1(e16, e10) = op1(e16, e14)) & ~ (op1(e16, e10) = op1(e16, e13)) & ~ (op1(e16, e10) = op1(e16, e12)) & ~ (op1(e16, e10) = op1(e16, e11)) & ~ (op1(e15, e15) = op1(e15, e16)) & ~ (op1(e15, e14) = op1(e15, e16)) & ~ (op1(e15, e14) = op1(e15, e15)) & ~ (op1(e15, e13) = op1(e15, e16)) & ~ (op1(e15, e13) = op1(e15, e15)) & ~ (op1(e15, e13) = op1(e15, e14)) & ~ (op1(e15, e12) = op1(e15, e16)) & ~ (op1(e15, e12) = op1(e15, e15)) & ~ (op1(e15, e12) = op1(e15, e14)) & ~ (op1(e15, e12) = op1(e15, e13)) & ~ (op1(e15, e11) = op1(e15, e16)) & ~ (op1(e15, e11) = op1(e15, e15)) & ~ (op1(e15, e11) = op1(e15, e14)) & ~ (op1(e15, e11) = op1(e15, e13)) & ~ (op1(e15, e11) = op1(e15, e12)) & ~ (op1(e15, e10) = op1(e15, e16)) & ~ (op1(e15, e10) = op1(e15, e15)) & ~ (op1(e15, e10) = op1(e15, e14)) & ~ (op1(e15, e10) = op1(e15, e13)) & ~ (op1(e15, e10) = op1(e15, e12)) & ~ (op1(e15, e10) = op1(e15, e11)) & ~ (op1(e14, e15) = op1(e14, e16)) & ~ (op1(e14, e14) = op1(e14, e16)) & ~ (op1(e14, e14) = op1(e14, e15)) & ~ (op1(e14, e13) = op1(e14, e16)) & ~ (op1(e14, e13) = op1(e14, e15)) & ~ (op1(e14, e13) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e16)) & ~ (op1(e14, e12) = op1(e14, e15)) & ~ (op1(e14, e12) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e16)) & ~ (op1(e14, e11) = op1(e14, e15)) & ~ (op1(e14, e11) = op1(e14, e14)) & ~ (op1(e14, e11) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e16)) & ~ (op1(e14, e10) = op1(e14, e15)) & ~ (op1(e14, e10) = op1(e14, e14)) & ~ (op1(e14, e10) = op1(e14, e13)) & ~ (op1(e14, e10) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e11)) & ~ (op1(e13, e15) = op1(e13, e16)) & ~ (op1(e13, e14) = op1(e13, e16)) & ~ (op1(e13, e14) = op1(e13, e15)) & ~ (op1(e13, e13) = op1(e13, e16)) & ~ (op1(e13, e13) = op1(e13, e15)) & ~ (op1(e13, e13) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e16)) & ~ (op1(e13, e12) = op1(e13, e15)) & ~ (op1(e13, e12) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e16)) & ~ (op1(e13, e11) = op1(e13, e15)) & ~ (op1(e13, e11) = op1(e13, e14)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e16)) & ~ (op1(e13, e10) = op1(e13, e15)) & ~ (op1(e13, e10) = op1(e13, e14)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e15) = op1(e12, e16)) & ~ (op1(e12, e14) = op1(e12, e16)) & ~ (op1(e12, e14) = op1(e12, e15)) & ~ (op1(e12, e13) = op1(e12, e16)) & ~ (op1(e12, e13) = op1(e12, e15)) & ~ (op1(e12, e13) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e16)) & ~ (op1(e12, e12) = op1(e12, e15)) & ~ (op1(e12, e12) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e16)) & ~ (op1(e12, e11) = op1(e12, e15)) & ~ (op1(e12, e11) = op1(e12, e14)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e16)) & ~ (op1(e12, e10) = op1(e12, e15)) & ~ (op1(e12, e10) = op1(e12, e14)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e15) = op1(e11, e16)) & ~ (op1(e11, e14) = op1(e11, e16)) & ~ (op1(e11, e14) = op1(e11, e15)) & ~ (op1(e11, e13) = op1(e11, e16)) & ~ (op1(e11, e13) = op1(e11, e15)) & ~ (op1(e11, e13) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e16)) & ~ (op1(e11, e12) = op1(e11, e15)) & ~ (op1(e11, e12) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e16)) & ~ (op1(e11, e11) = op1(e11, e15)) & ~ (op1(e11, e11) = op1(e11, e14)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e16)) & ~ (op1(e11, e10) = op1(e11, e15)) & ~ (op1(e11, e10) = op1(e11, e14)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e15) = op1(e10, e16)) & ~ (op1(e10, e14) = op1(e10, e16)) & ~ (op1(e10, e14) = op1(e10, e15)) & ~ (op1(e10, e13) = op1(e10, e16)) & ~ (op1(e10, e13) = op1(e10, e15)) & ~ (op1(e10, e13) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e16)) & ~ (op1(e10, e12) = op1(e10, e15)) & ~ (op1(e10, e12) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e16)) & ~ (op1(e10, e11) = op1(e10, e15)) & ~ (op1(e10, e11) = op1(e10, e14)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e16)) & ~ (op1(e10, e10) = op1(e10, e15)) & ~ (op1(e10, e10) = op1(e10, e14)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e15, e16) = op1(e16, e16)) & ~ (op1(e14, e16) = op1(e16, e16)) & ~ (op1(e14, e16) = op1(e15, e16)) & ~ (op1(e13, e16) = op1(e16, e16)) & ~ (op1(e13, e16) = op1(e15, e16)) & ~ (op1(e13, e16) = op1(e14, e16)) & ~ (op1(e12, e16) = op1(e16, e16)) & ~ (op1(e12, e16) = op1(e15, e16)) & ~ (op1(e12, e16) = op1(e14, e16)) & ~ (op1(e12, e16) = op1(e13, e16)) & ~ (op1(e11, e16) = op1(e16, e16)) & ~ (op1(e11, e16) = op1(e15, e16)) & ~ (op1(e11, e16) = op1(e14, e16)) & ~ (op1(e11, e16) = op1(e13, e16)) & ~ (op1(e11, e16) = op1(e12, e16)) & ~ (op1(e10, e16) = op1(e16, e16)) & ~ (op1(e10, e16) = op1(e15, e16)) & ~ (op1(e10, e16) = op1(e14, e16)) & ~ (op1(e10, e16) = op1(e13, e16)) & ~ (op1(e10, e16) = op1(e12, e16)) & ~ (op1(e10, e16) = op1(e11, e16)) & ~ (op1(e15, e15) = op1(e16, e15)) & ~ (op1(e14, e15) = op1(e16, e15)) & ~ (op1(e14, e15) = op1(e15, e15)) & ~ (op1(e13, e15) = op1(e16, e15)) & ~ (op1(e13, e15) = op1(e15, e15)) & ~ (op1(e13, e15) = op1(e14, e15)) & ~ (op1(e12, e15) = op1(e16, e15)) & ~ (op1(e12, e15) = op1(e15, e15)) & ~ (op1(e12, e15) = op1(e14, e15)) & ~ (op1(e12, e15) = op1(e13, e15)) & ~ (op1(e11, e15) = op1(e16, e15)) & ~ (op1(e11, e15) = op1(e15, e15)) & ~ (op1(e11, e15) = op1(e14, e15)) & ~ (op1(e11, e15) = op1(e13, e15)) & ~ (op1(e11, e15) = op1(e12, e15)) & ~ (op1(e10, e15) = op1(e16, e15)) & ~ (op1(e10, e15) = op1(e15, e15)) & ~ (op1(e10, e15) = op1(e14, e15)) & ~ (op1(e10, e15) = op1(e13, e15)) & ~ (op1(e10, e15) = op1(e12, e15)) & ~ (op1(e10, e15) = op1(e11, e15)) & ~ (op1(e15, e14) = op1(e16, e14)) & ~ (op1(e14, e14) = op1(e16, e14)) & ~ (op1(e14, e14) = op1(e15, e14)) & ~ (op1(e13, e14) = op1(e16, e14)) & ~ (op1(e13, e14) = op1(e15, e14)) & ~ (op1(e13, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e16, e14)) & ~ (op1(e12, e14) = op1(e15, e14)) & ~ (op1(e12, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e16, e14)) & ~ (op1(e11, e14) = op1(e15, e14)) & ~ (op1(e11, e14) = op1(e14, e14)) & ~ (op1(e11, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e16, e14)) & ~ (op1(e10, e14) = op1(e15, e14)) & ~ (op1(e10, e14) = op1(e14, e14)) & ~ (op1(e10, e14) = op1(e13, e14)) & ~ (op1(e10, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e11, e14)) & ~ (op1(e15, e13) = op1(e16, e13)) & ~ (op1(e14, e13) = op1(e16, e13)) & ~ (op1(e14, e13) = op1(e15, e13)) & ~ (op1(e13, e13) = op1(e16, e13)) & ~ (op1(e13, e13) = op1(e15, e13)) & ~ (op1(e13, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e16, e13)) & ~ (op1(e12, e13) = op1(e15, e13)) & ~ (op1(e12, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e16, e13)) & ~ (op1(e11, e13) = op1(e15, e13)) & ~ (op1(e11, e13) = op1(e14, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e16, e13)) & ~ (op1(e10, e13) = op1(e15, e13)) & ~ (op1(e10, e13) = op1(e14, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e15, e12) = op1(e16, e12)) & ~ (op1(e14, e12) = op1(e16, e12)) & ~ (op1(e14, e12) = op1(e15, e12)) & ~ (op1(e13, e12) = op1(e16, e12)) & ~ (op1(e13, e12) = op1(e15, e12)) & ~ (op1(e13, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e16, e12)) & ~ (op1(e12, e12) = op1(e15, e12)) & ~ (op1(e12, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e16, e12)) & ~ (op1(e11, e12) = op1(e15, e12)) & ~ (op1(e11, e12) = op1(e14, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e16, e12)) & ~ (op1(e10, e12) = op1(e15, e12)) & ~ (op1(e10, e12) = op1(e14, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e15, e11) = op1(e16, e11)) & ~ (op1(e14, e11) = op1(e16, e11)) & ~ (op1(e14, e11) = op1(e15, e11)) & ~ (op1(e13, e11) = op1(e16, e11)) & ~ (op1(e13, e11) = op1(e15, e11)) & ~ (op1(e13, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e16, e11)) & ~ (op1(e12, e11) = op1(e15, e11)) & ~ (op1(e12, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e16, e11)) & ~ (op1(e11, e11) = op1(e15, e11)) & ~ (op1(e11, e11) = op1(e14, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e16, e11)) & ~ (op1(e10, e11) = op1(e15, e11)) & ~ (op1(e10, e11) = op1(e14, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e15, e10) = op1(e16, e10)) & ~ (op1(e14, e10) = op1(e16, e10)) & ~ (op1(e14, e10) = op1(e15, e10)) & ~ (op1(e13, e10) = op1(e16, e10)) & ~ (op1(e13, e10) = op1(e15, e10)) & ~ (op1(e13, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e16, e10)) & ~ (op1(e12, e10) = op1(e15, e10)) & ~ (op1(e12, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e16, e10)) & ~ (op1(e11, e10) = op1(e15, e10)) & ~ (op1(e11, e10) = op1(e14, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e16, e10)) & ~ (op1(e10, e10) = op1(e15, e10)) & ~ (op1(e10, e10) = op1(e14, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax7)).
fof(f24496, plain, (spl42_70 | ~ spl42_100 | ~ spl42_265), inference(avatar_split_clause, [], [f24495, f2684, f2001, f1876])).
fof(f24495, plain, ((e16 = op1(e15, e14)) | (~ spl42_100 | ~ spl42_265)), inference(forward_demodulation, [], [f24488, f2686])).
fof(f24488, plain, ((e16 = op1(op1(e11, e14), e14)) | ~ spl42_100), inference(backward_demodulation, [], [f304, f2003])).
fof(f304, plain, (e16 = op1(op1(op1(e14, e16), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f24466, plain, (~ spl42_113 | ~ spl42_127), inference(avatar_split_clause, [], [f24455, f2113, f2055])).
fof(f2055, plain, (spl42_113 <=> (e10 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_113])])).
fof(f24455, plain, (~ (e10 = op1(e14, e14)) | ~ spl42_127), inference(backward_demodulation, [], [f746, f2115])).
fof(f746, plain, ~ (op1(e14, e12) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f24427, plain, (spl42_100 | ~ spl42_39 | ~ spl42_152), inference(avatar_split_clause, [], [f24426, f2216, f1748, f2001])).
fof(f24426, plain, ((e11 = op1(e14, e16)) | (~ spl42_39 | ~ spl42_152)), inference(backward_demodulation, [], [f12233, f2218])).
fof(f12233, plain, ((e11 = op1(op1(e13, e16), e16)) | ~ spl42_39), inference(backward_demodulation, [], [f271, f1750])).
fof(f271, plain, (e11 = op1(op1(op1(e16, e11), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f24385, plain, (spl42_122 | ~ spl42_26 | ~ spl42_182), inference(avatar_split_clause, [], [f24384, f2340, f1694, f2092])).
fof(f24384, plain, ((e12 = op1(e14, e13)) | (~ spl42_26 | ~ spl42_182)), inference(forward_demodulation, [], [f24377, f1696])).
fof(f24377, plain, ((e12 = op1(op1(e16, e13), e13)) | ~ spl42_182), inference(backward_demodulation, [], [f275, f2342])).
fof(f275, plain, (e12 = op1(op1(op1(e13, e12), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f24334, plain, (~ spl42_77 | ~ spl42_224), inference(avatar_split_clause, [], [f24330, f2514, f1905])).
fof(f1905, plain, (spl42_77 <=> (e16 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_77])])).
fof(f24330, plain, (~ (e16 = op1(e15, e13)) | ~ spl42_224), inference(backward_demodulation, [], [f579, f2516])).
fof(f579, plain, ~ (op1(e12, e13) = op1(e15, e13)), inference(cnf_transformation, [], [f7])).
fof(f24316, plain, (spl42_139 | ~ spl42_236 | ~ spl42_255), inference(avatar_split_clause, [], [f24315, f2643, f2564, f2162])).
fof(f24315, plain, ((e15 = op1(e14, e11)) | (~ spl42_236 | ~ spl42_255)), inference(backward_demodulation, [], [f13742, f2566])).
fof(f13742, plain, ((e15 = op1(op1(e12, e11), e11)) | ~ spl42_255), inference(forward_demodulation, [], [f294, f2645])).
fof(f294, plain, (e15 = op1(op1(op1(e11, e15), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f24305, plain, (spl42_127 | ~ spl42_240 | ~ spl42_278), inference(avatar_split_clause, [], [f24304, f2738, f2581, f2113])).
fof(f24304, plain, ((e10 = op1(e14, e12)) | (~ spl42_240 | ~ spl42_278)), inference(forward_demodulation, [], [f24290, f2740])).
fof(f24290, plain, ((e10 = op1(op1(e11, e12), e12)) | ~ spl42_240), inference(backward_demodulation, [], [f260, f2583])).
fof(f260, plain, (e10 = op1(op1(op1(e12, e10), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f24288, plain, (~ spl42_102 | ~ spl42_249), inference(avatar_split_clause, [], [f24282, f2618, f2009])).
fof(f2009, plain, (spl42_102 <=> (e13 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_102])])).
fof(f24282, plain, (~ (e13 = op1(e14, e16)) | ~ spl42_249), inference(backward_demodulation, [], [f637, f2620])).
fof(f637, plain, ~ (op1(e11, e16) = op1(e14, e16)), inference(cnf_transformation, [], [f7])).
fof(f24276, plain, (spl42_236 | ~ spl42_87 | ~ spl42_265), inference(avatar_split_clause, [], [f24275, f2684, f1947, f2564])).
fof(f24275, plain, ((e14 = op1(e12, e11)) | (~ spl42_87 | ~ spl42_265)), inference(forward_demodulation, [], [f24267, f1949])).
fof(f24267, plain, ((e14 = op1(op1(e15, e11), e11)) | ~ spl42_265), inference(backward_demodulation, [], [f287, f2686])).
fof(f287, plain, (e14 = op1(op1(op1(e11, e14), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f24256, plain, (~ spl42_190 | ~ spl42_288), inference(avatar_split_clause, [], [f24243, f2780, f2374])).
fof(f2374, plain, (spl42_190 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_190])])).
fof(f24243, plain, (~ (e10 = op1(e13, e10)) | ~ spl42_288), inference(backward_demodulation, [], [f510, f2782])).
fof(f510, plain, ~ (op1(e11, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f7])).
fof(f24240, plain, (spl42_112 | ~ spl42_50 | ~ spl42_306), inference(avatar_split_clause, [], [f24238, f2854, f1794, f2050])).
fof(f24238, plain, ((e16 = op1(e14, e15)) | (~ spl42_50 | ~ spl42_306)), inference(backward_demodulation, [], [f22687, f2856])).
fof(f22687, plain, ((e16 = op1(op1(e10, e15), e15)) | ~ spl42_50), inference(backward_demodulation, [], [f305, f1796])).
fof(f305, plain, (e16 = op1(op1(op1(e15, e16), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f24196, plain, (~ spl42_330 | ~ spl42_501 | ~ spl42_856 | ~ spl42_871 | spl42_1194), inference(avatar_contradiction_clause, [], [f24195])).
fof(f24195, plain, ($false | (~ spl42_330 | ~ spl42_501 | ~ spl42_856 | ~ spl42_871 | spl42_1194)), inference(subsumption_resolution, [], [f24189, f5757])).
fof(f24189, plain, (~ (e23 = h2(e10)) | (~ spl42_330 | ~ spl42_501 | ~ spl42_856 | ~ spl42_871 | spl42_1194)), inference(backward_demodulation, [], [f23919, f2956])).
fof(f23919, plain, (~ (e23 = h2(op1(e10, e11))) | (~ spl42_501 | ~ spl42_856 | ~ spl42_871 | spl42_1194)), inference(forward_demodulation, [], [f23909, f3762])).
fof(f3762, plain, ((e23 = op2(e23, e25)) | ~ spl42_501), inference(avatar_component_clause, [], [f3760])).
fof(f3760, plain, (spl42_501 <=> (e23 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_501])])).
fof(f23909, plain, (~ (op2(e23, e25) = h2(op1(e10, e11))) | (~ spl42_856 | ~ spl42_871 | spl42_1194)), inference(backward_demodulation, [], [f22967, f5757])).
fof(f22967, plain, (~ (h2(op1(e10, e11)) = op2(h2(e10), e25)) | (~ spl42_856 | spl42_1194)), inference(backward_demodulation, [], [f7888, f5682])).
fof(f7888, plain, (~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | spl42_1194), inference(avatar_component_clause, [], [f7886])).
fof(f7886, plain, (spl42_1194 <=> (h2(op1(e10, e11)) = op2(h2(e10), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1194])])).
fof(f24194, plain, (~ spl42_316 | ~ spl42_330), inference(avatar_split_clause, [], [f24180, f2954, f2896])).
fof(f2896, plain, (spl42_316 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_316])])).
fof(f24180, plain, (~ (e10 = op1(e10, e13)) | ~ spl42_330), inference(backward_demodulation, [], [f657, f2956])).
fof(f657, plain, ~ (op1(e10, e11) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f24144, plain, (spl42_413 | ~ spl42_443 | ~ spl42_608), inference(avatar_split_clause, [], [f24143, f4203, f3520, f3395])).
fof(f24143, plain, ((e26 = op2(e25, e24)) | (~ spl42_443 | ~ spl42_608)), inference(forward_demodulation, [], [f24136, f4205])).
fof(f24135, plain, (~ spl42_456 | ~ spl42_460), inference(avatar_contradiction_clause, [], [f24134])).
fof(f24134, plain, ($false | (~ spl42_456 | ~ spl42_460)), inference(subsumption_resolution, [], [f24133, f1115])).
fof(f1115, plain, ~ (e20 = e24), inference(cnf_transformation, [], [f10])).
fof(f24133, plain, ((e20 = e24) | (~ spl42_456 | ~ spl42_460)), inference(backward_demodulation, [], [f3592, f3576])).
fof(f3576, plain, ((e20 = op2(e24, e24)) | ~ spl42_456), inference(avatar_component_clause, [], [f3574])).
fof(f3574, plain, (spl42_456 <=> (e20 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_456])])).
fof(f24098, plain, (spl42_443 | ~ spl42_495 | ~ spl42_698), inference(avatar_split_clause, [], [f24095, f4872, f3735, f3520])).
fof(f24095, plain, ((e21 = op2(e24, e26)) | (~ spl42_495 | ~ spl42_698)), inference(backward_demodulation, [], [f15201, f3737])).
fof(f24043, plain, (~ spl42_374 | ~ spl42_567 | ~ spl42_568), inference(avatar_contradiction_clause, [], [f24042])).
fof(f24042, plain, ($false | (~ spl42_374 | ~ spl42_567 | ~ spl42_568)), inference(subsumption_resolution, [], [f24041, f1114])).
fof(f1114, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f10])).
fof(f24041, plain, ((e20 = e23) | (~ spl42_374 | ~ spl42_567 | ~ spl42_568)), inference(forward_demodulation, [], [f24040, f4040])).
fof(f4040, plain, ((e20 = op2(e22, e22)) | ~ spl42_568), inference(avatar_component_clause, [], [f4038])).
fof(f4038, plain, (spl42_568 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_568])])).
fof(f24012, plain, (spl42_860 | ~ spl42_621 | ~ spl42_875), inference(avatar_split_clause, [], [f24010, f5776, f4257, f5701])).
fof(f4257, plain, (spl42_621 <=> (e24 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_621])])).
fof(f24010, plain, ((e24 = h2(e14)) | (~ spl42_621 | ~ spl42_875)), inference(backward_demodulation, [], [f23895, f4259])).
fof(f4259, plain, ((e24 = op2(e21, e22)) | ~ spl42_621), inference(avatar_component_clause, [], [f4257])).
fof(f23968, plain, (spl42_487 | ~ spl42_439 | ~ spl42_664), inference(avatar_split_clause, [], [f23967, f4435, f3503, f3702])).
fof(f23967, plain, ((e23 = op2(e24, e20)) | (~ spl42_439 | ~ spl42_664)), inference(forward_demodulation, [], [f23960, f3505])).
fof(f23960, plain, ((e23 = op2(op2(e25, e20), e20)) | ~ spl42_664), inference(backward_demodulation, [], [f475, f4437])).
fof(f475, plain, (e23 = op2(op2(op2(e20, e23), e20), e20)), inference(cnf_transformation, [], [f6])).
fof(f23952, plain, (~ spl42_652 | ~ spl42_673), inference(avatar_split_clause, [], [f23942, f4473, f4386])).
fof(f4386, plain, (spl42_652 <=> (e20 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_652])])).
fof(f23942, plain, (~ (e20 = op2(e20, e24)) | ~ spl42_673), inference(backward_demodulation, [], [f952, f4475])).
fof(f952, plain, ~ (op2(e20, e21) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f23862, plain, (~ spl42_535 | ~ spl42_908), inference(avatar_split_clause, [], [f23846, f5942, f3901])).
fof(f3901, plain, (spl42_535 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_535])])).
fof(f5942, plain, (spl42_908 <=> (e22 = h1(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_908])])).
fof(f23846, plain, (~ (e22 = op2(e23, e20)) | ~ spl42_908), inference(backward_demodulation, [], [f4647, f5943])).
fof(f5943, plain, ((e22 = h1(e15)) | ~ spl42_908), inference(avatar_component_clause, [], [f5942])).
fof(f4647, plain, ~ (op2(e23, e20) = h1(e15)), inference(backward_demodulation, [], [f799, f1204])).
fof(f1204, plain, (op2(e20, e20) = h1(e15)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(e20, e20) = h1(e15)) & (h1(e14) = op2(e20, op2(e20, op2(op2(e20, e20), op2(e20, e20))))) & (h1(e13) = op2(e20, op2(op2(e20, e20), op2(e20, e20)))) & (h1(e12) = op2(op2(op2(e20, e20), op2(e20, e20)), op2(e20, e20))) & (h1(e11) = op2(op2(e20, e20), op2(e20, e20))) & (h1(e10) = op2(e20, op2(e20, e20))) & (e20 = h1(e16))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax16)).
fof(f799, plain, ~ (op2(e20, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f8])).
fof(f23842, plain, (spl42_381 | ~ spl42_729 | ~ spl42_733 | ~ spl42_1227), inference(avatar_contradiction_clause, [], [f23841])).
fof(f23841, plain, ($false | (spl42_381 | ~ spl42_729 | ~ spl42_733 | ~ spl42_1227)), inference(subsumption_resolution, [], [f23838, f3264])).
fof(f3264, plain, (~ (e22 = op2(e26, e21)) | spl42_381), inference(avatar_component_clause, [], [f3263])).
fof(f3263, plain, (spl42_381 <=> (e22 = op2(e26, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_381])])).
fof(f23838, plain, ((e22 = op2(e26, e21)) | (~ spl42_729 | ~ spl42_733 | ~ spl42_1227)), inference(backward_demodulation, [], [f23834, f8019])).
fof(f23834, plain, ((e22 = op2(h2(e15), e21)) | (~ spl42_729 | ~ spl42_733)), inference(forward_demodulation, [], [f14048, f5044])).
fof(f5044, plain, ((e22 = h6(e12)) | ~ spl42_729), inference(avatar_component_clause, [], [f5043])).
fof(f5043, plain, (spl42_729 <=> (e22 = h6(e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_729])])).
fof(f14048, plain, ((h6(e12) = op2(h2(e15), e21)) | ~ spl42_733), inference(backward_demodulation, [], [f13003, f14046])).
fof(f14046, plain, ((h2(e15) = h6(e11)) | ~ spl42_733), inference(backward_demodulation, [], [f13002, f1211])).
fof(f13002, plain, ((op2(e21, e21) = h6(e11)) | ~ spl42_733), inference(backward_demodulation, [], [f4766, f5064])).
fof(f5064, plain, ((e21 = h6(e15)) | ~ spl42_733), inference(avatar_component_clause, [], [f5063])).
fof(f5063, plain, (spl42_733 <=> (e21 = h6(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_733])])).
fof(f4766, plain, (h6(e11) = op2(h6(e15), h6(e15))), inference(forward_demodulation, [], [f1235, f1239])).
fof(f1239, plain, (op2(e25, e25) = h6(e15)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ((op2(e25, e25) = h6(e15)) & (h6(e14) = op2(e25, op2(e25, op2(op2(e25, e25), op2(e25, e25))))) & (h6(e13) = op2(e25, op2(op2(e25, e25), op2(e25, e25)))) & (h6(e12) = op2(op2(op2(e25, e25), op2(e25, e25)), op2(e25, e25))) & (h6(e11) = op2(op2(e25, e25), op2(e25, e25))) & (h6(e10) = op2(e25, op2(e25, e25))) & (e25 = h6(e16))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax21)).
fof(f1235, plain, (h6(e11) = op2(op2(e25, e25), op2(e25, e25))), inference(cnf_transformation, [], [f21])).
fof(f13003, plain, ((h6(e12) = op2(h6(e11), e21)) | ~ spl42_733), inference(backward_demodulation, [], [f4767, f5064])).
fof(f4767, plain, (h6(e12) = op2(h6(e11), h6(e15))), inference(backward_demodulation, [], [f4765, f4766])).
fof(f4765, plain, (h6(e12) = op2(op2(h6(e15), h6(e15)), h6(e15))), inference(forward_demodulation, [], [f1236, f1239])).
fof(f1236, plain, (h6(e12) = op2(op2(op2(e25, e25), op2(e25, e25)), op2(e25, e25))), inference(cnf_transformation, [], [f21])).
fof(f23800, plain, (~ spl42_637 | ~ spl42_392), inference(avatar_split_clause, [], [f23799, f3308, f4323])).
fof(f4323, plain, (spl42_637 <=> (e26 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_637])])).
fof(f23799, plain, (~ (e26 = op2(e21, e20)) | ~ spl42_392), inference(forward_demodulation, [], [f807, f3310])).
fof(f807, plain, ~ (op2(e21, e20) = op2(e26, e20)), inference(cnf_transformation, [], [f8])).
fof(f23762, plain, (~ spl42_476 | ~ spl42_455), inference(avatar_split_clause, [], [f23761, f3569, f3656])).
fof(f3656, plain, (spl42_476 <=> (e26 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_476])])).
fof(f23761, plain, (~ (e26 = op2(e24, e22)) | ~ spl42_455), inference(forward_demodulation, [], [f1041, f3571])).
fof(f1041, plain, ~ (op2(e24, e22) = op2(e24, e25)), inference(cnf_transformation, [], [f8])).
fof(f23759, plain, (~ spl42_442 | ~ spl42_393), inference(avatar_split_clause, [], [f23758, f3313, f3516])).
fof(f3516, plain, (spl42_442 <=> (e20 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_442])])).
fof(f3313, plain, (spl42_393 <=> (e20 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_393])])).
fof(f23758, plain, (~ (e20 = op2(e24, e26)) | ~ spl42_393), inference(forward_demodulation, [], [f941, f3315])).
fof(f3315, plain, ((e20 = op2(e25, e26)) | ~ spl42_393), inference(avatar_component_clause, [], [f3313])).
fof(f941, plain, ~ (op2(e24, e26) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f23738, plain, (~ spl42_227 | ~ spl42_31), inference(avatar_split_clause, [], [f23737, f1715, f2527])).
fof(f2527, plain, (spl42_227 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_227])])).
fof(f23737, plain, (~ (e12 = op1(e12, e12)) | ~ spl42_31), inference(forward_demodulation, [], [f559, f1717])).
fof(f559, plain, ~ (op1(e12, e12) = op1(e16, e12)), inference(cnf_transformation, [], [f7])).
fof(f23713, plain, (~ spl42_16 | ~ spl42_140 | spl42_261), inference(avatar_contradiction_clause, [], [f23712])).
fof(f23712, plain, ($false | (~ spl42_16 | ~ spl42_140 | spl42_261)), inference(subsumption_resolution, [], [f23707, f2669])).
fof(f2669, plain, (~ (e11 = op1(e11, e14)) | spl42_261), inference(avatar_component_clause, [], [f2668])).
fof(f2668, plain, (spl42_261 <=> (e11 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_261])])).
fof(f23707, plain, ((e11 = op1(e11, e14)) | (~ spl42_16 | ~ spl42_140)), inference(backward_demodulation, [], [f21495, f1655])).
fof(f21495, plain, ((e11 = op1(op1(e16, e14), e14)) | ~ spl42_140), inference(backward_demodulation, [], [f269, f2168])).
fof(f2168, plain, ((e16 = op1(e14, e11)) | ~ spl42_140), inference(avatar_component_clause, [], [f2166])).
fof(f2166, plain, (spl42_140 <=> (e16 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_140])])).
fof(f269, plain, (e11 = op1(op1(op1(e14, e11), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f23689, plain, (~ spl42_26 | ~ spl42_102 | ~ spl42_150), inference(avatar_contradiction_clause, [], [f23688])).
fof(f23688, plain, ($false | (~ spl42_26 | ~ spl42_102 | ~ spl42_150)), inference(subsumption_resolution, [], [f23687, f1102])).
fof(f1102, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f9])).
fof(f23687, plain, ((e12 = e13) | (~ spl42_26 | ~ spl42_102 | ~ spl42_150)), inference(forward_demodulation, [], [f23682, f2210])).
fof(f2210, plain, ((e12 = op1(e13, e16)) | ~ spl42_150), inference(avatar_component_clause, [], [f2208])).
fof(f2208, plain, (spl42_150 <=> (e12 = op1(e13, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_150])])).
fof(f23682, plain, ((e13 = op1(e13, e16)) | (~ spl42_26 | ~ spl42_102)), inference(backward_demodulation, [], [f12209, f2011])).
fof(f2011, plain, ((e13 = op1(e14, e16)) | ~ spl42_102), inference(avatar_component_clause, [], [f2009])).
fof(f12209, plain, ((e13 = op1(op1(e14, e16), e16)) | ~ spl42_26), inference(backward_demodulation, [], [f285, f1696])).
fof(f285, plain, (e13 = op1(op1(op1(e16, e13), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f23530, plain, (~ spl42_103 | ~ spl42_250), inference(avatar_split_clause, [], [f23523, f2622, f2013])).
fof(f2013, plain, (spl42_103 <=> (e14 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_103])])).
fof(f2622, plain, (spl42_250 <=> (e14 = op1(e11, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_250])])).
fof(f23523, plain, (~ (e14 = op1(e14, e16)) | ~ spl42_250), inference(backward_demodulation, [], [f637, f2624])).
fof(f2624, plain, ((e14 = op1(e11, e16)) | ~ spl42_250), inference(avatar_component_clause, [], [f2622])).
fof(f23471, plain, (~ spl42_280 | ~ spl42_287), inference(avatar_split_clause, [], [f23466, f2775, f2746])).
fof(f2746, plain, (spl42_280 <=> (e16 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_280])])).
fof(f23466, plain, (~ (e16 = op1(e11, e12)) | ~ spl42_287), inference(backward_demodulation, [], [f677, f2777])).
fof(f677, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f7])).
fof(f23460, plain, (spl42_232 | ~ spl42_87 | ~ spl42_293), inference(avatar_split_clause, [], [f23459, f2800, f1947, f2548])).
fof(f2548, plain, (spl42_232 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_232])])).
fof(f2800, plain, (spl42_293 <=> (e15 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_293])])).
fof(f23459, plain, ((e10 = op1(e12, e11)) | (~ spl42_87 | ~ spl42_293)), inference(forward_demodulation, [], [f23448, f1949])).
fof(f23448, plain, ((e10 = op1(op1(e15, e11), e11)) | ~ spl42_293), inference(backward_demodulation, [], [f259, f2802])).
fof(f2802, plain, ((e15 = op1(e11, e10)) | ~ spl42_293), inference(avatar_component_clause, [], [f2800])).
fof(f259, plain, (e10 = op1(op1(op1(e11, e10), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f23387, plain, (~ spl42_364 | ~ spl42_392), inference(avatar_split_clause, [], [f23382, f3308, f3192])).
fof(f23382, plain, (~ (e26 = op2(e26, e24)) | ~ spl42_392), inference(backward_demodulation, [], [f1073, f3310])).
fof(f1073, plain, ~ (op2(e26, e20) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f23379, plain, (spl42_883 | ~ spl42_393 | ~ spl42_856 | ~ spl42_1227), inference(avatar_split_clause, [], [f23377, f8018, f5681, f3313, f5817])).
fof(f23377, plain, ((e20 = h2(e12)) | (~ spl42_393 | ~ spl42_856 | ~ spl42_1227)), inference(backward_demodulation, [], [f22965, f3315])).
fof(f22965, plain, ((op2(e25, e26) = h2(e12)) | (~ spl42_856 | ~ spl42_1227)), inference(backward_demodulation, [], [f22934, f5682])).
fof(f22934, plain, ((h2(e12) = op2(h2(e11), e26)) | ~ spl42_1227), inference(backward_demodulation, [], [f4683, f8019])).
fof(f4683, plain, (h2(e12) = op2(h2(e11), h2(e15))), inference(backward_demodulation, [], [f4681, f4682])).
fof(f4682, plain, (h2(e11) = op2(h2(e15), h2(e15))), inference(forward_demodulation, [], [f1207, f1211])).
fof(f1207, plain, (h2(e11) = op2(op2(e21, e21), op2(e21, e21))), inference(cnf_transformation, [], [f17])).
fof(f4681, plain, (h2(e12) = op2(op2(h2(e15), h2(e15)), h2(e15))), inference(forward_demodulation, [], [f1208, f1211])).
fof(f1208, plain, (h2(e12) = op2(op2(op2(e21, e21), op2(e21, e21)), op2(e21, e21))), inference(cnf_transformation, [], [f17])).
fof(f23369, plain, (~ spl42_401 | ~ spl42_404), inference(avatar_contradiction_clause, [], [f23368])).
fof(f23368, plain, ($false | (~ spl42_401 | ~ spl42_404)), inference(subsumption_resolution, [], [f23367, f1120])).
fof(f23367, plain, ((e21 = e24) | (~ spl42_401 | ~ spl42_404)), inference(forward_demodulation, [], [f3360, f3348])).
fof(f3360, plain, ((e24 = op2(e25, e25)) | ~ spl42_404), inference(avatar_component_clause, [], [f3358])).
fof(f3358, plain, (spl42_404 <=> (e24 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_404])])).
fof(f23356, plain, (~ spl42_417 | ~ spl42_420), inference(avatar_contradiction_clause, [], [f23355])).
fof(f23355, plain, ($false | (~ spl42_417 | ~ spl42_420)), inference(subsumption_resolution, [], [f23354, f1129])).
fof(f1129, plain, ~ (e23 = e26), inference(cnf_transformation, [], [f10])).
fof(f23354, plain, ((e23 = e26) | (~ spl42_417 | ~ spl42_420)), inference(forward_demodulation, [], [f3426, f3414])).
fof(f23252, plain, (~ spl42_510 | ~ spl42_531), inference(avatar_split_clause, [], [f23244, f3884, f3797])).
fof(f3797, plain, (spl42_510 <=> (e25 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_510])])).
fof(f3884, plain, (spl42_531 <=> (e25 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_531])])).
fof(f23244, plain, (~ (e25 = op2(e23, e24)) | ~ spl42_531), inference(backward_demodulation, [], [f1015, f3886])).
fof(f3886, plain, ((e25 = op2(e23, e21)) | ~ spl42_531), inference(avatar_component_clause, [], [f3884])).
fof(f1015, plain, ~ (op2(e23, e21) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f23251, plain, (~ spl42_482 | ~ spl42_531), inference(avatar_split_clause, [], [f23242, f3884, f3681])).
fof(f23242, plain, (~ (e25 = op2(e24, e21)) | ~ spl42_531), inference(backward_demodulation, [], [f833, f3886])).
fof(f833, plain, ~ (op2(e23, e21) = op2(e24, e21)), inference(cnf_transformation, [], [f8])).
fof(f23192, plain, (spl42_482 | ~ spl42_579 | ~ spl42_598), inference(avatar_split_clause, [], [f23190, f4162, f4083, f3681])).
fof(f4162, plain, (spl42_598 <=> (e22 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_598])])).
fof(f23190, plain, ((e25 = op2(e24, e21)) | (~ spl42_579 | ~ spl42_598)), inference(backward_demodulation, [], [f14033, f4085])).
fof(f14033, plain, ((e25 = op2(op2(e22, e21), e21)) | ~ spl42_598), inference(forward_demodulation, [], [f490, f4164])).
fof(f4164, plain, ((e22 = op2(e21, e25)) | ~ spl42_598), inference(avatar_component_clause, [], [f4162])).
fof(f490, plain, (e25 = op2(op2(op2(e21, e25), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f23191, plain, (~ spl42_544 | ~ spl42_579), inference(avatar_split_clause, [], [f23187, f4083, f3938])).
fof(f3938, plain, (spl42_544 <=> (e24 = op2(e22, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_544])])).
fof(f23187, plain, (~ (e24 = op2(e22, e26)) | ~ spl42_579), inference(backward_demodulation, [], [f996, f4085])).
fof(f996, plain, ~ (op2(e22, e21) = op2(e22, e26)), inference(cnf_transformation, [], [f8])).
fof(f23150, plain, (~ spl42_624 | ~ spl42_630), inference(avatar_contradiction_clause, [], [f23149])).
fof(f23149, plain, ($false | (~ spl42_624 | ~ spl42_630)), inference(subsumption_resolution, [], [f23144, f1117])).
fof(f1117, plain, ~ (e20 = e26), inference(cnf_transformation, [], [f10])).
fof(f23144, plain, ((e20 = e26) | (~ spl42_624 | ~ spl42_630)), inference(backward_demodulation, [], [f4296, f4272])).
fof(f4272, plain, ((e20 = op2(e21, e21)) | ~ spl42_624), inference(avatar_component_clause, [], [f4270])).
fof(f4270, plain, (spl42_624 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_624])])).
fof(f22977, plain, (spl42_875 | ~ spl42_598 | ~ spl42_856), inference(avatar_split_clause, [], [f22976, f5681, f4162, f5776])).
fof(f22976, plain, ((e22 = h2(e13)) | (~ spl42_598 | ~ spl42_856)), inference(forward_demodulation, [], [f22968, f4164])).
fof(f22968, plain, ((op2(e21, e25) = h2(e13)) | ~ spl42_856), inference(backward_demodulation, [], [f4684, f5682])).
fof(f4684, plain, (h2(e13) = op2(e21, h2(e11))), inference(backward_demodulation, [], [f4679, f4682])).
fof(f22943, plain, (spl42_856 | ~ spl42_349 | ~ spl42_1227), inference(avatar_split_clause, [], [f22942, f8018, f3130, f5681])).
fof(f3130, plain, (spl42_349 <=> (e25 = op2(e26, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_349])])).
fof(f22942, plain, ((e25 = h2(e11)) | (~ spl42_349 | ~ spl42_1227)), inference(forward_demodulation, [], [f22933, f3132])).
fof(f3132, plain, ((e25 = op2(e26, e26)) | ~ spl42_349), inference(avatar_component_clause, [], [f3130])).
fof(f22933, plain, ((op2(e26, e26) = h2(e11)) | ~ spl42_1227), inference(backward_demodulation, [], [f4682, f8019])).
fof(f22899, plain, (~ spl42_675 | ~ spl42_731 | ~ spl42_733), inference(avatar_split_clause, [], [f22898, f5063, f5053, f4481])).
fof(f4481, plain, (spl42_675 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_675])])).
fof(f5053, plain, (spl42_731 <=> (e22 = h6(e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_731])])).
fof(f22898, plain, (~ (e22 = op2(e20, e21)) | (~ spl42_731 | ~ spl42_733)), inference(forward_demodulation, [], [f18840, f5054])).
fof(f5054, plain, ((e22 = h6(e10)) | ~ spl42_731), inference(avatar_component_clause, [], [f5053])).
fof(f18840, plain, (~ (op2(e20, e21) = h6(e10)) | ~ spl42_733), inference(forward_demodulation, [], [f822, f13004])).
fof(f13004, plain, ((op2(e25, e21) = h6(e10)) | ~ spl42_733), inference(backward_demodulation, [], [f4769, f5064])).
fof(f4769, plain, (h6(e10) = op2(e25, h6(e15))), inference(forward_demodulation, [], [f1234, f1239])).
fof(f1234, plain, (h6(e10) = op2(e25, op2(e25, e25))), inference(cnf_transformation, [], [f21])).
fof(f822, plain, ~ (op2(e20, e21) = op2(e25, e21)), inference(cnf_transformation, [], [f8])).
fof(f22895, plain, (~ spl42_656 | ~ spl42_460), inference(avatar_split_clause, [], [f22894, f3590, f4402])).
fof(f4402, plain, (spl42_656 <=> (e24 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_656])])).
fof(f22894, plain, (~ (e24 = op2(e20, e24)) | ~ spl42_460), inference(forward_demodulation, [], [f4729, f21851])).
fof(f21851, plain, ((e24 = h5(e15)) | ~ spl42_460), inference(forward_demodulation, [], [f1232, f3592])).
fof(f4729, plain, ~ (op2(e20, e24) = h5(e15)), inference(backward_demodulation, [], [f884, f1232])).
fof(f884, plain, ~ (op2(e20, e24) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f22832, plain, (spl42_578 | ~ spl42_430 | ~ spl42_615), inference(avatar_split_clause, [], [f22251, f4232, f3466, f4079])).
fof(f4079, plain, (spl42_578 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_578])])).
fof(f4232, plain, (spl42_615 <=> (e25 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_615])])).
fof(f22251, plain, ((e23 = op2(e22, e21)) | (~ spl42_430 | ~ spl42_615)), inference(backward_demodulation, [], [f22064, f3468])).
fof(f22064, plain, ((e23 = op2(op2(e25, e21), e21)) | ~ spl42_615), inference(backward_demodulation, [], [f476, f4234])).
fof(f4234, plain, ((e25 = op2(e21, e23)) | ~ spl42_615), inference(avatar_component_clause, [], [f4232])).
fof(f476, plain, (e23 = op2(op2(op2(e21, e23), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f22803, plain, (spl42_590 | ~ spl42_492 | ~ spl42_698), inference(avatar_split_clause, [], [f22802, f4872, f3723, f4129])).
fof(f4129, plain, (spl42_590 <=> (e21 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_590])])).
fof(f3723, plain, (spl42_492 <=> (e21 = op2(e23, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_492])])).
fof(f22802, plain, ((e21 = op2(e21, e26)) | (~ spl42_492 | ~ spl42_698)), inference(forward_demodulation, [], [f15201, f3725])).
fof(f3725, plain, ((e21 = op2(e23, e26)) | ~ spl42_492), inference(avatar_component_clause, [], [f3723])).
fof(f22675, plain, (~ spl42_70 | ~ spl42_77), inference(avatar_split_clause, [], [f22671, f1905, f1876])).
fof(f22671, plain, (~ (e16 = op1(e15, e14)) | ~ spl42_77), inference(backward_demodulation, [], [f770, f1907])).
fof(f1907, plain, ((e16 = op1(e15, e13)) | ~ spl42_77), inference(avatar_component_clause, [], [f1905])).
fof(f770, plain, ~ (op1(e15, e13) = op1(e15, e14)), inference(cnf_transformation, [], [f7])).
fof(f22656, plain, (~ spl42_74 | ~ spl42_95), inference(avatar_split_clause, [], [f22649, f1980, f1893])).
fof(f1980, plain, (spl42_95 <=> (e13 = op1(e15, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_95])])).
fof(f22649, plain, (~ (e13 = op1(e15, e13)) | ~ spl42_95), inference(backward_demodulation, [], [f757, f1982])).
fof(f1982, plain, ((e13 = op1(e15, e10)) | ~ spl42_95), inference(avatar_component_clause, [], [f1980])).
fof(f757, plain, ~ (op1(e15, e10) = op1(e15, e13)), inference(cnf_transformation, [], [f7])).
fof(f22645, plain, (~ spl42_26 | spl42_102 | ~ spl42_103), inference(avatar_contradiction_clause, [], [f22644])).
fof(f22644, plain, ($false | (~ spl42_26 | spl42_102 | ~ spl42_103)), inference(subsumption_resolution, [], [f22639, f2010])).
fof(f2010, plain, (~ (e13 = op1(e14, e16)) | spl42_102), inference(avatar_component_clause, [], [f2009])).
fof(f22639, plain, ((e13 = op1(e14, e16)) | (~ spl42_26 | ~ spl42_103)), inference(backward_demodulation, [], [f12209, f2015])).
fof(f2015, plain, ((e14 = op1(e14, e16)) | ~ spl42_103), inference(avatar_component_clause, [], [f2013])).
fof(f22588, plain, (~ spl42_74 | ~ spl42_172), inference(avatar_split_clause, [], [f22580, f2299, f1893])).
fof(f2299, plain, (spl42_172 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_172])])).
fof(f22580, plain, (~ (e13 = op1(e15, e13)) | ~ spl42_172), inference(backward_demodulation, [], [f582, f2301])).
fof(f2301, plain, ((e13 = op1(e13, e13)) | ~ spl42_172), inference(avatar_component_clause, [], [f2299])).
fof(f582, plain, ~ (op1(e13, e13) = op1(e15, e13)), inference(cnf_transformation, [], [f7])).
fof(f22564, plain, (~ spl42_183 | ~ spl42_184), inference(avatar_contradiction_clause, [], [f22563])).
fof(f22563, plain, ($false | (~ spl42_183 | ~ spl42_184)), inference(subsumption_resolution, [], [f22562, f1091])).
fof(f1091, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f9])).
fof(f22562, plain, ((e10 = e11) | (~ spl42_183 | ~ spl42_184)), inference(backward_demodulation, [], [f2351, f2347])).
fof(f2347, plain, ((e10 = op1(e13, e11)) | ~ spl42_183), inference(avatar_component_clause, [], [f2345])).
fof(f2345, plain, (spl42_183 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_183])])).
fof(f22420, plain, (spl42_137 | ~ spl42_267 | ~ spl42_334), inference(avatar_contradiction_clause, [], [f22419])).
fof(f22419, plain, ($false | (spl42_137 | ~ spl42_267 | ~ spl42_334)), inference(subsumption_resolution, [], [f22418, f2155])).
fof(f2155, plain, (~ (e13 = op1(e14, e11)) | spl42_137), inference(avatar_component_clause, [], [f2154])).
fof(f2154, plain, (spl42_137 <=> (e13 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_137])])).
fof(f22418, plain, ((e13 = op1(e14, e11)) | (~ spl42_267 | ~ spl42_334)), inference(forward_demodulation, [], [f22406, f2972])).
fof(f2972, plain, ((e14 = op1(e10, e11)) | ~ spl42_334), inference(avatar_component_clause, [], [f2970])).
fof(f2970, plain, (spl42_334 <=> (e14 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_334])])).
fof(f22406, plain, ((e13 = op1(op1(e10, e11), e11)) | ~ spl42_267), inference(backward_demodulation, [], [f280, f2695])).
fof(f2695, plain, ((e10 = op1(e11, e13)) | ~ spl42_267), inference(avatar_component_clause, [], [f2693])).
fof(f2693, plain, (spl42_267 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_267])])).
fof(f280, plain, (e13 = op1(op1(op1(e11, e13), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f22404, plain, (~ spl42_39 | spl42_185 | ~ spl42_280), inference(avatar_contradiction_clause, [], [f22403])).
fof(f22403, plain, ($false | (~ spl42_39 | spl42_185 | ~ spl42_280)), inference(subsumption_resolution, [], [f22402, f2354])).
fof(f2354, plain, (~ (e12 = op1(e13, e11)) | spl42_185), inference(avatar_component_clause, [], [f2353])).
fof(f22402, plain, ((e12 = op1(e13, e11)) | (~ spl42_39 | ~ spl42_280)), inference(forward_demodulation, [], [f22395, f1750])).
fof(f22395, plain, ((e12 = op1(op1(e16, e11), e11)) | ~ spl42_280), inference(backward_demodulation, [], [f273, f2748])).
fof(f2748, plain, ((e16 = op1(e11, e12)) | ~ spl42_280), inference(avatar_component_clause, [], [f2746])).
fof(f273, plain, (e12 = op1(op1(op1(e11, e12), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f22347, plain, (~ spl42_306 | ~ spl42_327), inference(avatar_split_clause, [], [f22339, f2941, f2854])).
fof(f2941, plain, (spl42_327 <=> (e14 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_327])])).
fof(f22339, plain, (~ (e14 = op1(e10, e15)) | ~ spl42_327), inference(backward_demodulation, [], [f663, f2943])).
fof(f2943, plain, ((e14 = op1(e10, e12)) | ~ spl42_327), inference(avatar_component_clause, [], [f2941])).
fof(f663, plain, ~ (op1(e10, e12) = op1(e10, e15)), inference(cnf_transformation, [], [f7])).
fof(f22324, plain, (~ spl42_236 | ~ spl42_334), inference(avatar_split_clause, [], [f22314, f2970, f2564])).
fof(f22314, plain, (~ (e14 = op1(e12, e11)) | ~ spl42_334), inference(backward_demodulation, [], [f525, f2972])).
fof(f525, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f7])).
fof(f22297, plain, (~ spl42_359 | ~ spl42_373), inference(avatar_split_clause, [], [f22289, f3230, f3172])).
fof(f3230, plain, (spl42_373 <=> (e21 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_373])])).
fof(f22289, plain, (~ (e21 = op2(e26, e24)) | ~ spl42_373), inference(backward_demodulation, [], [f1082, f3232])).
fof(f3232, plain, ((e21 = op2(e26, e22)) | ~ spl42_373), inference(avatar_component_clause, [], [f3230])).
fof(f1082, plain, ~ (op2(e26, e22) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f22284, plain, (~ spl42_374 | ~ spl42_388), inference(avatar_split_clause, [], [f22274, f3292, f3234])).
fof(f22274, plain, (~ (e22 = op2(e26, e22)) | ~ spl42_388), inference(backward_demodulation, [], [f1071, f3294])).
fof(f1071, plain, ~ (op2(e26, e20) = op2(e26, e22)), inference(cnf_transformation, [], [f8])).
fof(f22283, plain, (~ spl42_388 | spl42_442 | ~ spl42_544), inference(avatar_contradiction_clause, [], [f22282])).
fof(f22282, plain, ($false | (~ spl42_388 | spl42_442 | ~ spl42_544)), inference(subsumption_resolution, [], [f22281, f3517])).
fof(f3517, plain, (~ (e20 = op2(e24, e26)) | spl42_442), inference(avatar_component_clause, [], [f3516])).
fof(f22281, plain, ((e20 = op2(e24, e26)) | (~ spl42_388 | ~ spl42_544)), inference(forward_demodulation, [], [f22273, f3940])).
fof(f3940, plain, ((e24 = op2(e22, e26)) | ~ spl42_544), inference(avatar_component_clause, [], [f3938])).
fof(f22273, plain, ((e20 = op2(op2(e22, e26), e26)) | ~ spl42_388), inference(backward_demodulation, [], [f460, f3294])).
fof(f22246, plain, (~ spl42_439 | spl42_485 | ~ spl42_678), inference(avatar_contradiction_clause, [], [f22245])).
fof(f22245, plain, ($false | (~ spl42_439 | spl42_485 | ~ spl42_678)), inference(subsumption_resolution, [], [f22241, f3695])).
fof(f3695, plain, (~ (e21 = op2(e24, e20)) | spl42_485), inference(avatar_component_clause, [], [f3694])).
fof(f3694, plain, (spl42_485 <=> (e21 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_485])])).
fof(f22241, plain, ((e21 = op2(e24, e20)) | (~ spl42_439 | ~ spl42_678)), inference(backward_demodulation, [], [f20958, f3505])).
fof(f20958, plain, ((e21 = op2(op2(e25, e20), e20)) | ~ spl42_678), inference(backward_demodulation, [], [f461, f4495])).
fof(f4495, plain, ((e25 = op2(e20, e21)) | ~ spl42_678), inference(avatar_component_clause, [], [f4493])).
fof(f4493, plain, (spl42_678 <=> (e25 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_678])])).
fof(f461, plain, (e21 = op2(op2(op2(e20, e21), e20), e20)), inference(cnf_transformation, [], [f6])).
fof(f22244, plain, (~ spl42_439 | spl42_449 | ~ spl42_453), inference(avatar_contradiction_clause, [], [f22243])).
fof(f22243, plain, ($false | (~ spl42_439 | spl42_449 | ~ spl42_453)), inference(subsumption_resolution, [], [f22242, f3546])).
fof(f3546, plain, (~ (e20 = op2(e24, e25)) | spl42_449), inference(avatar_component_clause, [], [f3545])).
fof(f3545, plain, (spl42_449 <=> (e20 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_449])])).
fof(f22242, plain, ((e20 = op2(e24, e25)) | (~ spl42_439 | ~ spl42_453)), inference(forward_demodulation, [], [f22236, f3563])).
fof(f22236, plain, ((e20 = op2(op2(e24, e25), e25)) | ~ spl42_439), inference(backward_demodulation, [], [f459, f3505])).
fof(f459, plain, (e20 = op2(op2(op2(e25, e20), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f22188, plain, (spl42_345 | ~ spl42_497 | ~ spl42_698), inference(avatar_contradiction_clause, [], [f22187])).
fof(f22187, plain, ($false | (spl42_345 | ~ spl42_497 | ~ spl42_698)), inference(subsumption_resolution, [], [f22183, f3115])).
fof(f3115, plain, (~ (e21 = op2(e26, e26)) | spl42_345), inference(avatar_component_clause, [], [f3114])).
fof(f3114, plain, (spl42_345 <=> (e21 = op2(e26, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_345])])).
fof(f22183, plain, ((e21 = op2(e26, e26)) | (~ spl42_497 | ~ spl42_698)), inference(backward_demodulation, [], [f15201, f3745])).
fof(f3745, plain, ((e26 = op2(e23, e26)) | ~ spl42_497), inference(avatar_component_clause, [], [f3743])).
fof(f3743, plain, (spl42_497 <=> (e26 = op2(e23, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_497])])).
fof(f22178, plain, (~ spl42_501 | ~ spl42_514 | spl42_566), inference(avatar_contradiction_clause, [], [f22177])).
fof(f22177, plain, ($false | (~ spl42_501 | ~ spl42_514 | spl42_566)), inference(subsumption_resolution, [], [f22176, f4030])).
fof(f4030, plain, (~ (e25 = op2(e22, e23)) | spl42_566), inference(avatar_component_clause, [], [f4029])).
fof(f4029, plain, (spl42_566 <=> (e25 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_566])])).
fof(f22176, plain, ((e25 = op2(e22, e23)) | (~ spl42_501 | ~ spl42_514)), inference(forward_demodulation, [], [f22169, f3816])).
fof(f3816, plain, ((e22 = op2(e23, e23)) | ~ spl42_514), inference(avatar_component_clause, [], [f3814])).
fof(f3814, plain, (spl42_514 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_514])])).
fof(f22169, plain, ((e25 = op2(op2(e23, e23), e23)) | ~ spl42_501), inference(backward_demodulation, [], [f492, f3762])).
fof(f22119, plain, (~ spl42_374 | ~ spl42_570), inference(avatar_split_clause, [], [f22106, f4046, f3234])).
fof(f4046, plain, (spl42_570 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_570])])).
fof(f22106, plain, (~ (e22 = op2(e26, e22)) | ~ spl42_570), inference(backward_demodulation, [], [f4692, f22101])).
fof(f22101, plain, ((e22 = h3(e15)) | ~ spl42_570), inference(backward_demodulation, [], [f1218, f4048])).
fof(f4048, plain, ((e22 = op2(e22, e22)) | ~ spl42_570), inference(avatar_component_clause, [], [f4046])).
fof(f4692, plain, ~ (op2(e26, e22) = h3(e15)), inference(backward_demodulation, [], [f853, f1218])).
fof(f853, plain, ~ (op2(e22, e22) = op2(e26, e22)), inference(cnf_transformation, [], [f8])).
fof(f22082, plain, (~ spl42_393 | ~ spl42_589), inference(avatar_split_clause, [], [f22071, f4125, f3313])).
fof(f4125, plain, (spl42_589 <=> (e20 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_589])])).
fof(f22071, plain, (~ (e20 = op2(e25, e26)) | ~ spl42_589), inference(backward_demodulation, [], [f932, f4127])).
fof(f4127, plain, ((e20 = op2(e21, e26)) | ~ spl42_589), inference(avatar_component_clause, [], [f4125])).
fof(f932, plain, ~ (op2(e21, e26) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f22063, plain, (~ spl42_592 | ~ spl42_620), inference(avatar_split_clause, [], [f22058, f4253, f4137])).
fof(f4253, plain, (spl42_620 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_620])])).
fof(f22058, plain, (~ (e23 = op2(e21, e26)) | ~ spl42_620), inference(backward_demodulation, [], [f979, f4255])).
fof(f4255, plain, ((e23 = op2(e21, e22)) | ~ spl42_620), inference(avatar_component_clause, [], [f4253])).
fof(f979, plain, ~ (op2(e21, e22) = op2(e21, e26)), inference(cnf_transformation, [], [f8])).
fof(f22029, plain, (~ spl42_644 | ~ spl42_651), inference(avatar_split_clause, [], [f22024, f4381, f4352])).
fof(f4381, plain, (spl42_651 <=> (e26 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_651])])).
fof(f22024, plain, (~ (e26 = op2(e20, e26)) | ~ spl42_651), inference(backward_demodulation, [], [f964, f4383])).
fof(f4383, plain, ((e26 = op2(e20, e25)) | ~ spl42_651), inference(avatar_component_clause, [], [f4381])).
fof(f964, plain, ~ (op2(e20, e25) = op2(e20, e26)), inference(cnf_transformation, [], [f8])).
fof(f22017, plain, (~ spl42_359 | ~ spl42_653), inference(avatar_split_clause, [], [f22005, f4390, f3172])).
fof(f4390, plain, (spl42_653 <=> (e21 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_653])])).
fof(f22005, plain, (~ (e21 = op2(e26, e24)) | ~ spl42_653), inference(backward_demodulation, [], [f886, f4392])).
fof(f4392, plain, ((e21 = op2(e20, e24)) | ~ spl42_653), inference(avatar_component_clause, [], [f4390])).
fof(f886, plain, ~ (op2(e20, e24) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f21976, plain, (~ spl42_677 | ~ spl42_678), inference(avatar_contradiction_clause, [], [f21975])).
fof(f21975, plain, ($false | (~ spl42_677 | ~ spl42_678)), inference(subsumption_resolution, [], [f21974, f1130])).
fof(f1130, plain, ~ (e24 = e25), inference(cnf_transformation, [], [f10])).
fof(f21974, plain, ((e24 = e25) | (~ spl42_677 | ~ spl42_678)), inference(backward_demodulation, [], [f4495, f4491])).
fof(f4491, plain, ((e24 = op2(e20, e21)) | ~ spl42_677), inference(avatar_component_clause, [], [f4489])).
fof(f4489, plain, (spl42_677 <=> (e24 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_677])])).
fof(f21844, plain, (spl42_359 | ~ spl42_413 | ~ spl42_482), inference(avatar_split_clause, [], [f21843, f3681, f3395, f3172])).
fof(f21843, plain, ((e21 = op2(e26, e24)) | (~ spl42_413 | ~ spl42_482)), inference(forward_demodulation, [], [f21842, f3397])).
fof(f21842, plain, ((e21 = op2(op2(e25, e24), e24)) | ~ spl42_482), inference(forward_demodulation, [], [f465, f3683])).
fof(f465, plain, (e21 = op2(op2(op2(e24, e21), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f21825, plain, (~ spl42_671 | ~ spl42_426), inference(avatar_split_clause, [], [f20862, f3449, f4464])).
fof(f4464, plain, (spl42_671 <=> (e25 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_671])])).
fof(f20862, plain, (~ (e25 = op2(e20, e22)) | ~ spl42_426), inference(forward_demodulation, [], [f843, f3451])).
fof(f843, plain, ~ (op2(e20, e22) = op2(e25, e22)), inference(cnf_transformation, [], [f8])).
fof(f21814, plain, (spl42_649 | ~ spl42_351 | ~ spl42_413), inference(avatar_split_clause, [], [f21193, f3395, f3139, f4373])).
fof(f21193, plain, ((e24 = op2(e20, e25)) | (~ spl42_351 | ~ spl42_413)), inference(forward_demodulation, [], [f21192, f3141])).
fof(f21192, plain, ((e24 = op2(op2(e26, e25), e25)) | ~ spl42_413), inference(backward_demodulation, [], [f487, f3397])).
fof(f21790, plain, (~ spl42_613 | ~ spl42_417), inference(avatar_split_clause, [], [f21789, f3412, f4224])).
fof(f4224, plain, (spl42_613 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_613])])).
fof(f21789, plain, (~ (e23 = op2(e21, e23)) | ~ spl42_417), inference(forward_demodulation, [], [f869, f3414])).
fof(f869, plain, ~ (op2(e21, e23) = op2(e25, e23)), inference(cnf_transformation, [], [f8])).
fof(f21754, plain, (~ spl42_364 | ~ spl42_413), inference(avatar_split_clause, [], [f21753, f3395, f3192])).
fof(f21753, plain, (~ (e26 = op2(e26, e24)) | ~ spl42_413), inference(forward_demodulation, [], [f901, f3397])).
fof(f901, plain, ~ (op2(e25, e24) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f21731, plain, (~ spl42_336 | ~ spl42_301), inference(avatar_split_clause, [], [f21730, f2833, f2978])).
fof(f2978, plain, (spl42_336 <=> (e16 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_336])])).
fof(f21730, plain, (~ (e16 = op1(e10, e11)) | ~ spl42_301), inference(forward_demodulation, [], [f660, f2835])).
fof(f660, plain, ~ (op1(e10, e11) = op1(e10, e16)), inference(cnf_transformation, [], [f7])).
fof(f21707, plain, (~ spl42_252 | ~ spl42_301), inference(avatar_split_clause, [], [f21706, f2833, f2630])).
fof(f2630, plain, (spl42_252 <=> (e16 = op1(e11, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_252])])).
fof(f21706, plain, (~ (e16 = op1(e11, e16)) | ~ spl42_301), inference(forward_demodulation, [], [f629, f2835])).
fof(f629, plain, ~ (op1(e10, e16) = op1(e11, e16)), inference(cnf_transformation, [], [f7])).
fof(f21643, plain, (~ spl42_56 | ~ spl42_301), inference(avatar_split_clause, [], [f21642, f2833, f1818])).
fof(f1818, plain, (spl42_56 <=> (e16 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_56])])).
fof(f21642, plain, (~ (e16 = op1(e15, e16)) | ~ spl42_301), inference(forward_demodulation, [], [f633, f2835])).
fof(f633, plain, ~ (op1(e10, e16) = op1(e15, e16)), inference(cnf_transformation, [], [f7])).
fof(f21633, plain, (~ spl42_17 | ~ spl42_199), inference(avatar_contradiction_clause, [], [f21632])).
fof(f21632, plain, ($false | (~ spl42_17 | ~ spl42_199)), inference(subsumption_resolution, [], [f21631, f1103])).
fof(f1103, plain, ~ (e12 = e14), inference(cnf_transformation, [], [f9])).
fof(f21631, plain, ((e12 = e14) | (~ spl42_17 | ~ spl42_199)), inference(forward_demodulation, [], [f21630, f2413])).
fof(f21630, plain, ((e14 = op1(e12, e16)) | (~ spl42_17 | ~ spl42_199)), inference(forward_demodulation, [], [f21624, f2413])).
fof(f21624, plain, ((e14 = op1(op1(e12, e16), e16)) | ~ spl42_17), inference(backward_demodulation, [], [f292, f1659])).
fof(f1659, plain, ((e12 = op1(e16, e14)) | ~ spl42_17), inference(avatar_component_clause, [], [f1657])).
fof(f1657, plain, (spl42_17 <=> (e12 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_17])])).
fof(f292, plain, (e14 = op1(op1(op1(e16, e14), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f21623, plain, (~ spl42_26 | ~ spl42_28), inference(avatar_contradiction_clause, [], [f21622])).
fof(f21622, plain, ($false | (~ spl42_26 | ~ spl42_28)), inference(subsumption_resolution, [], [f21621, f1110])).
fof(f1110, plain, ~ (e14 = e16), inference(cnf_transformation, [], [f9])).
fof(f21621, plain, ((e14 = e16) | (~ spl42_26 | ~ spl42_28)), inference(forward_demodulation, [], [f1704, f1696])).
fof(f1704, plain, ((e16 = op1(e16, e13)) | ~ spl42_28), inference(avatar_component_clause, [], [f1702])).
fof(f1702, plain, (spl42_28 <=> (e16 = op1(e16, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_28])])).
fof(f21595, plain, (~ spl42_21 | ~ spl42_70), inference(avatar_split_clause, [], [f21593, f1876, f1673])).
fof(f1673, plain, (spl42_21 <=> (e16 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_21])])).
fof(f21593, plain, (~ (e16 = op1(e16, e14)) | ~ spl42_70), inference(backward_demodulation, [], [f607, f1878])).
fof(f607, plain, ~ (op1(e15, e14) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f21563, plain, (spl42_57 | ~ spl42_94 | ~ spl42_209), inference(avatar_contradiction_clause, [], [f21562])).
fof(f21562, plain, ($false | (spl42_57 | ~ spl42_94 | ~ spl42_209)), inference(subsumption_resolution, [], [f21561, f1824])).
fof(f1824, plain, (~ (e10 = op1(e15, e15)) | spl42_57), inference(avatar_component_clause, [], [f1823])).
fof(f1823, plain, (spl42_57 <=> (e10 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_57])])).
fof(f21561, plain, ((e10 = op1(e15, e15)) | (~ spl42_94 | ~ spl42_209)), inference(forward_demodulation, [], [f21551, f2454])).
fof(f21551, plain, ((e10 = op1(op1(e12, e15), e15)) | ~ spl42_94), inference(backward_demodulation, [], [f263, f1978])).
fof(f1978, plain, ((e12 = op1(e15, e10)) | ~ spl42_94), inference(avatar_component_clause, [], [f1976])).
fof(f1976, plain, (spl42_94 <=> (e12 = op1(e15, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_94])])).
fof(f263, plain, (e10 = op1(op1(op1(e15, e10), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f21505, plain, (~ spl42_133 | ~ spl42_140), inference(avatar_split_clause, [], [f21497, f2166, f2137])).
fof(f2137, plain, (spl42_133 <=> (e16 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_133])])).
fof(f21497, plain, (~ (e16 = op1(e14, e12)) | ~ spl42_140), inference(backward_demodulation, [], [f740, f2168])).
fof(f740, plain, ~ (op1(e14, e11) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f21481, plain, (~ spl42_39 | ~ spl42_149 | spl42_247), inference(avatar_contradiction_clause, [], [f21480])).
fof(f21480, plain, ($false | (~ spl42_39 | ~ spl42_149 | spl42_247)), inference(subsumption_resolution, [], [f21475, f2611])).
fof(f2611, plain, (~ (e11 = op1(e11, e16)) | spl42_247), inference(avatar_component_clause, [], [f2610])).
fof(f2610, plain, (spl42_247 <=> (e11 = op1(e11, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_247])])).
fof(f21475, plain, ((e11 = op1(e11, e16)) | (~ spl42_39 | ~ spl42_149)), inference(backward_demodulation, [], [f12233, f2206])).
fof(f2206, plain, ((e11 = op1(e13, e16)) | ~ spl42_149), inference(avatar_component_clause, [], [f2204])).
fof(f2204, plain, (spl42_149 <=> (e11 = op1(e13, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_149])])).
fof(f21465, plain, (~ spl42_112 | ~ spl42_161), inference(avatar_split_clause, [], [f21458, f2253, f2050])).
fof(f2253, plain, (spl42_161 <=> (e16 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_161])])).
fof(f21458, plain, (~ (e16 = op1(e14, e15)) | ~ spl42_161), inference(backward_demodulation, [], [f623, f2255])).
fof(f2255, plain, ((e16 = op1(e13, e15)) | ~ spl42_161), inference(avatar_component_clause, [], [f2253])).
fof(f623, plain, ~ (op1(e13, e15) = op1(e14, e15)), inference(cnf_transformation, [], [f7])).
fof(f21456, plain, (spl42_75 | ~ spl42_162 | ~ spl42_321), inference(avatar_contradiction_clause, [], [f21455])).
fof(f21455, plain, ($false | (spl42_75 | ~ spl42_162 | ~ spl42_321)), inference(subsumption_resolution, [], [f21454, f1898])).
fof(f1898, plain, (~ (e14 = op1(e15, e13)) | spl42_75), inference(avatar_component_clause, [], [f1897])).
fof(f1897, plain, (spl42_75 <=> (e14 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_75])])).
fof(f21454, plain, ((e14 = op1(e15, e13)) | (~ spl42_162 | ~ spl42_321)), inference(forward_demodulation, [], [f21442, f2918])).
fof(f21442, plain, ((e14 = op1(op1(e10, e13), e13)) | ~ spl42_162), inference(backward_demodulation, [], [f289, f2260])).
fof(f2260, plain, ((e10 = op1(e13, e14)) | ~ spl42_162), inference(avatar_component_clause, [], [f2258])).
fof(f2258, plain, (spl42_162 <=> (e10 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_162])])).
fof(f289, plain, (e14 = op1(op1(op1(e13, e14), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f21411, plain, (spl42_35 | ~ spl42_199 | ~ spl42_231), inference(avatar_contradiction_clause, [], [f21410])).
fof(f21410, plain, ($false | (spl42_35 | ~ spl42_199 | ~ spl42_231)), inference(subsumption_resolution, [], [f21409, f1732])).
fof(f1732, plain, (~ (e16 = op1(e16, e12)) | spl42_35), inference(avatar_component_clause, [], [f1731])).
fof(f1731, plain, (spl42_35 <=> (e16 = op1(e16, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_35])])).
fof(f21409, plain, ((e16 = op1(e16, e12)) | (~ spl42_199 | ~ spl42_231)), inference(forward_demodulation, [], [f21402, f2545])).
fof(f2545, plain, ((e16 = op1(e12, e12)) | ~ spl42_231), inference(avatar_component_clause, [], [f2543])).
fof(f2543, plain, (spl42_231 <=> (e16 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_231])])).
fof(f21402, plain, ((e16 = op1(op1(e12, e12), e12)) | ~ spl42_199), inference(backward_demodulation, [], [f302, f2413])).
fof(f302, plain, (e16 = op1(op1(op1(e12, e16), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f21318, plain, (spl42_233 | ~ spl42_281 | ~ spl42_332), inference(avatar_split_clause, [], [f21317, f2962, f2751, f2552])).
fof(f2552, plain, (spl42_233 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_233])])).
fof(f2751, plain, (spl42_281 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_281])])).
fof(f2962, plain, (spl42_332 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_332])])).
fof(f21317, plain, ((e11 = op1(e12, e11)) | (~ spl42_281 | ~ spl42_332)), inference(forward_demodulation, [], [f21304, f2964])).
fof(f2964, plain, ((e12 = op1(e10, e11)) | ~ spl42_332), inference(avatar_component_clause, [], [f2962])).
fof(f21304, plain, ((e11 = op1(op1(e10, e11), e11)) | ~ spl42_281), inference(backward_demodulation, [], [f266, f2753])).
fof(f2753, plain, ((e10 = op1(e11, e11)) | ~ spl42_281), inference(avatar_component_clause, [], [f2751])).
fof(f266, plain, (e11 = op1(op1(op1(e11, e11), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f21245, plain, (~ spl42_185 | ~ spl42_332), inference(avatar_split_clause, [], [f21235, f2962, f2353])).
fof(f21235, plain, (~ (e12 = op1(e13, e11)) | ~ spl42_332), inference(backward_demodulation, [], [f526, f2964])).
fof(f526, plain, ~ (op1(e10, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f21199, plain, (~ spl42_401 | ~ spl42_402), inference(avatar_contradiction_clause, [], [f21198])).
fof(f21198, plain, ($false | (~ spl42_401 | ~ spl42_402)), inference(subsumption_resolution, [], [f21197, f1118])).
fof(f1118, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f10])).
fof(f21197, plain, ((e21 = e22) | (~ spl42_401 | ~ spl42_402)), inference(forward_demodulation, [], [f3352, f3348])).
fof(f3352, plain, ((e22 = op2(e25, e25)) | ~ spl42_402), inference(avatar_component_clause, [], [f3350])).
fof(f3350, plain, (spl42_402 <=> (e22 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_402])])).
fof(f21191, plain, (~ spl42_417 | spl42_515 | ~ spl42_782), inference(avatar_contradiction_clause, [], [f21190])).
fof(f21190, plain, ($false | (~ spl42_417 | spl42_515 | ~ spl42_782)), inference(subsumption_resolution, [], [f21184, f3819])).
fof(f3819, plain, (~ (e23 = op2(e23, e23)) | spl42_515), inference(avatar_component_clause, [], [f3818])).
fof(f3818, plain, (spl42_515 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_515])])).
fof(f21184, plain, ((e23 = op2(e23, e23)) | (~ spl42_417 | ~ spl42_782)), inference(backward_demodulation, [], [f20933, f3414])).
fof(f20933, plain, ((e23 = op2(op2(e25, e23), e23)) | ~ spl42_782), inference(backward_demodulation, [], [f4707, f5310])).
fof(f5310, plain, ((e25 = h4(e15)) | ~ spl42_782), inference(avatar_component_clause, [], [f5309])).
fof(f5309, plain, (spl42_782 <=> (e25 = h4(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_782])])).
fof(f4707, plain, (e23 = op2(op2(h4(e15), e23), e23)), inference(backward_demodulation, [], [f478, f1225])).
fof(f478, plain, (e23 = op2(op2(op2(e23, e23), e23), e23)), inference(cnf_transformation, [], [f6])).
fof(f20995, plain, (spl42_455 | ~ spl42_393 | ~ spl42_649), inference(avatar_split_clause, [], [f20993, f4373, f3313, f3569])).
fof(f20993, plain, ((e26 = op2(e24, e25)) | (~ spl42_393 | ~ spl42_649)), inference(backward_demodulation, [], [f20100, f4375])).
fof(f20100, plain, ((e26 = op2(op2(e20, e25), e25)) | ~ spl42_393), inference(backward_demodulation, [], [f501, f3315])).
fof(f20898, plain, (spl42_742 | ~ spl42_393 | ~ spl42_733 | ~ spl42_1227), inference(avatar_split_clause, [], [f20897, f8018, f5063, f3313, f5108])).
fof(f20897, plain, ((e20 = h6(e13)) | (~ spl42_393 | ~ spl42_733 | ~ spl42_1227)), inference(forward_demodulation, [], [f20885, f3315])).
fof(f20885, plain, ((op2(e25, e26) = h6(e13)) | (~ spl42_733 | ~ spl42_1227)), inference(backward_demodulation, [], [f14049, f8019])).
fof(f14049, plain, ((h6(e13) = op2(e25, h2(e15))) | ~ spl42_733), inference(backward_demodulation, [], [f4768, f14046])).
fof(f4768, plain, (h6(e13) = op2(e25, h6(e11))), inference(backward_demodulation, [], [f4763, f4766])).
fof(f4763, plain, (h6(e13) = op2(e25, op2(h6(e15), h6(e15)))), inference(forward_demodulation, [], [f1237, f1239])).
fof(f1237, plain, (h6(e13) = op2(e25, op2(op2(e25, e25), op2(e25, e25)))), inference(cnf_transformation, [], [f21])).
fof(f20893, plain, (spl42_871 | ~ spl42_592 | ~ spl42_1227), inference(avatar_split_clause, [], [f20892, f8018, f4137, f5756])).
fof(f20892, plain, ((e23 = h2(e10)) | (~ spl42_592 | ~ spl42_1227)), inference(forward_demodulation, [], [f20882, f4139])).
fof(f20882, plain, ((op2(e21, e26) = h2(e10)) | ~ spl42_1227), inference(backward_demodulation, [], [f4685, f8019])).
fof(f4685, plain, (h2(e10) = op2(e21, h2(e15))), inference(forward_demodulation, [], [f1206, f1211])).
fof(f1206, plain, (h2(e10) = op2(e21, op2(e21, e21))), inference(cnf_transformation, [], [f17])).
fof(f20889, plain, (~ spl42_483 | ~ spl42_1227), inference(avatar_split_clause, [], [f20878, f8018, f3685])).
fof(f20878, plain, (~ (e26 = op2(e24, e21)) | ~ spl42_1227), inference(backward_demodulation, [], [f4669, f8019])).
fof(f4669, plain, ~ (op2(e24, e21) = h2(e15)), inference(backward_demodulation, [], [f826, f1211])).
fof(f826, plain, ~ (op2(e21, e21) = op2(e24, e21)), inference(cnf_transformation, [], [f8])).
fof(f20843, plain, (spl42_418 | ~ spl42_505 | ~ spl42_664), inference(avatar_contradiction_clause, [], [f20842])).
fof(f20842, plain, ($false | (spl42_418 | ~ spl42_505 | ~ spl42_664)), inference(subsumption_resolution, [], [f20841, f3417])).
fof(f3417, plain, (~ (e24 = op2(e25, e23)) | spl42_418), inference(avatar_component_clause, [], [f3416])).
fof(f3416, plain, (spl42_418 <=> (e24 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_418])])).
fof(f20841, plain, ((e24 = op2(e25, e23)) | (~ spl42_505 | ~ spl42_664)), inference(forward_demodulation, [], [f19103, f4437])).
fof(f19103, plain, ((e24 = op2(op2(e20, e23), e23)) | ~ spl42_505), inference(backward_demodulation, [], [f485, f3779])).
fof(f3779, plain, ((e20 = op2(e23, e24)) | ~ spl42_505), inference(avatar_component_clause, [], [f3777])).
fof(f3777, plain, (spl42_505 <=> (e20 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_505])])).
fof(f485, plain, (e24 = op2(op2(op2(e23, e24), e23), e23)), inference(cnf_transformation, [], [f6])).
fof(f20830, plain, (~ spl42_646 | ~ spl42_733), inference(avatar_split_clause, [], [f15195, f5063, f4361])).
fof(f4361, plain, (spl42_646 <=> (e21 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_646])])).
fof(f15195, plain, (~ (e21 = op2(e20, e25)) | ~ spl42_733), inference(forward_demodulation, [], [f4750, f5064])).
fof(f4750, plain, ~ (op2(e20, e25) = h6(e15)), inference(backward_demodulation, [], [f906, f1239])).
fof(f906, plain, ~ (op2(e20, e25) = op2(e25, e25)), inference(cnf_transformation, [], [f8])).
fof(f20829, plain, (~ spl42_679 | ~ spl42_644), inference(avatar_split_clause, [], [f20828, f4352, f4497])).
fof(f4497, plain, (spl42_679 <=> (e26 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_679])])).
fof(f20828, plain, (~ (e26 = op2(e20, e21)) | ~ spl42_644), inference(forward_demodulation, [], [f954, f4354])).
fof(f954, plain, ~ (op2(e20, e21) = op2(e20, e26)), inference(cnf_transformation, [], [f8])).
fof(f20823, plain, (~ spl42_658 | ~ spl42_644), inference(avatar_split_clause, [], [f20822, f4352, f4410])).
fof(f4410, plain, (spl42_658 <=> (e26 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_658])])).
fof(f20822, plain, (~ (e26 = op2(e20, e24)) | ~ spl42_644), inference(forward_demodulation, [], [f963, f4354])).
fof(f963, plain, ~ (op2(e20, e24) = op2(e20, e26)), inference(cnf_transformation, [], [f8])).
fof(f20811, plain, (spl42_644 | ~ spl42_393 | ~ spl42_688), inference(avatar_split_clause, [], [f20107, f4816, f3313, f4352])).
fof(f20107, plain, ((e26 = op2(e20, e26)) | (~ spl42_393 | ~ spl42_688)), inference(backward_demodulation, [], [f12687, f3315])).
fof(f20807, plain, (~ spl42_606 | ~ spl42_592), inference(avatar_split_clause, [], [f20806, f4137, f4195])).
fof(f4195, plain, (spl42_606 <=> (e23 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_606])])).
fof(f20806, plain, (~ (e23 = op2(e21, e24)) | ~ spl42_592), inference(forward_demodulation, [], [f984, f4139])).
fof(f984, plain, ~ (op2(e21, e24) = op2(e21, e26)), inference(cnf_transformation, [], [f8])).
fof(f20797, plain, (~ spl42_831 | ~ spl42_669), inference(avatar_split_clause, [], [f20796, f4456, f5555])).
fof(f5555, plain, (spl42_831 <=> (e23 = h3(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_831])])).
fof(f4456, plain, (spl42_669 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_669])])).
fof(f20796, plain, (~ (e23 = h3(e15)) | ~ spl42_669), inference(forward_demodulation, [], [f4687, f4458])).
fof(f4458, plain, ((e23 = op2(e20, e22)) | ~ spl42_669), inference(avatar_component_clause, [], [f4456])).
fof(f4687, plain, ~ (op2(e20, e22) = h3(e15)), inference(backward_demodulation, [], [f840, f1218])).
fof(f840, plain, ~ (op2(e20, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f8])).
fof(f20790, plain, (~ spl42_559 | ~ spl42_552), inference(avatar_split_clause, [], [f16015, f3971, f4000])).
fof(f4000, plain, (spl42_559 <=> (e25 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_559])])).
fof(f16015, plain, (~ (e25 = op2(e22, e24)) | ~ spl42_552), inference(forward_demodulation, [], [f1004, f3973])).
fof(f1004, plain, ~ (op2(e22, e24) = op2(e22, e25)), inference(cnf_transformation, [], [f8])).
fof(f20783, plain, (~ spl42_543 | ~ spl42_592), inference(avatar_split_clause, [], [f20782, f4137, f3934])).
fof(f3934, plain, (spl42_543 <=> (e23 = op2(e22, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_543])])).
fof(f20782, plain, (~ (e23 = op2(e22, e26)) | ~ spl42_592), inference(forward_demodulation, [], [f929, f4139])).
fof(f929, plain, ~ (op2(e21, e26) = op2(e22, e26)), inference(cnf_transformation, [], [f8])).
fof(f20756, plain, (~ spl42_1117 | ~ spl42_567), inference(avatar_split_clause, [], [f20755, f4033, f7322])).
fof(f7322, plain, (spl42_1117 <=> (e26 = h4(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_1117])])).
fof(f20755, plain, (~ (e26 = h4(e15)) | ~ spl42_567), inference(forward_demodulation, [], [f4710, f4035])).
fof(f4710, plain, ~ (op2(e22, e23) = h4(e15)), inference(backward_demodulation, [], [f871, f1225])).
fof(f871, plain, ~ (op2(e22, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f20737, plain, (~ spl42_436 | ~ spl42_733), inference(avatar_split_clause, [], [f13907, f5063, f3491])).
fof(f3491, plain, (spl42_436 <=> (e21 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_436])])).
fof(f13907, plain, (~ (e21 = op2(e25, e20)) | ~ spl42_733), inference(forward_demodulation, [], [f4756, f5064])).
fof(f4756, plain, ~ (op2(e25, e20) = h6(e15)), inference(backward_demodulation, [], [f1053, f1239])).
fof(f1053, plain, ~ (op2(e25, e20) = op2(e25, e25)), inference(cnf_transformation, [], [f8])).
fof(f20730, plain, (~ spl42_387 | ~ spl42_359), inference(avatar_split_clause, [], [f20729, f3172, f3288])).
fof(f3288, plain, (spl42_387 <=> (e21 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_387])])).
fof(f20729, plain, (~ (e21 = op2(e26, e20)) | ~ spl42_359), inference(forward_demodulation, [], [f1073, f3174])).
fof(f20712, plain, (spl42_298 | ~ spl42_26 | ~ spl42_99), inference(avatar_split_clause, [], [f20495, f1997, f1694, f2821])).
fof(f2821, plain, (spl42_298 <=> (e13 = op1(e10, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_298])])).
fof(f1997, plain, (spl42_99 <=> (e10 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_99])])).
fof(f20495, plain, ((e13 = op1(e10, e16)) | (~ spl42_26 | ~ spl42_99)), inference(backward_demodulation, [], [f12209, f1999])).
fof(f1999, plain, ((e10 = op1(e14, e16)) | ~ spl42_99), inference(avatar_component_clause, [], [f1997])).
fof(f20662, plain, (spl42_154 | ~ spl42_6 | ~ spl42_53), inference(avatar_split_clause, [], [f20577, f1806, f1611, f2224])).
fof(f2224, plain, (spl42_154 <=> (e16 = op1(e13, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_154])])).
fof(f1806, plain, (spl42_53 <=> (e13 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_53])])).
fof(f20577, plain, ((e16 = op1(e13, e16)) | (~ spl42_6 | ~ spl42_53)), inference(backward_demodulation, [], [f12249, f1808])).
fof(f1808, plain, ((e13 = op1(e15, e16)) | ~ spl42_53), inference(avatar_component_clause, [], [f1806])).
fof(f12249, plain, ((e16 = op1(op1(e15, e16), e16)) | ~ spl42_6), inference(forward_demodulation, [], [f306, f1613])).
fof(f306, plain, (e16 = op1(op1(op1(e16, e16), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f20572, plain, (~ spl42_58 | ~ spl42_59), inference(avatar_contradiction_clause, [], [f20571])).
fof(f20571, plain, ($false | (~ spl42_58 | ~ spl42_59)), inference(subsumption_resolution, [], [f20570, f1097])).
fof(f20570, plain, ((e11 = e12) | (~ spl42_58 | ~ spl42_59)), inference(forward_demodulation, [], [f1833, f1829])).
fof(f1833, plain, ((e12 = op1(e15, e15)) | ~ spl42_59), inference(avatar_component_clause, [], [f1831])).
fof(f1831, plain, (spl42_59 <=> (e12 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_59])])).
fof(f20553, plain, (~ spl42_67 | ~ spl42_74), inference(avatar_split_clause, [], [f20543, f1893, f1864])).
fof(f1864, plain, (spl42_67 <=> (e13 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_67])])).
fof(f20543, plain, (~ (e13 = op1(e15, e14)) | ~ spl42_74), inference(backward_demodulation, [], [f770, f1895])).
fof(f20481, plain, (~ spl42_112 | ~ spl42_133), inference(avatar_contradiction_clause, [], [f20480])).
fof(f20480, plain, ($false | (~ spl42_112 | ~ spl42_133)), inference(subsumption_resolution, [], [f20479, f1104])).
fof(f1104, plain, ~ (e12 = e15), inference(cnf_transformation, [], [f9])).
fof(f20479, plain, ((e12 = e15) | (~ spl42_112 | ~ spl42_133)), inference(backward_demodulation, [], [f19480, f20475])).
fof(f20475, plain, ((e12 = op1(op1(e16, e14), e14)) | ~ spl42_133), inference(backward_demodulation, [], [f276, f2139])).
fof(f2139, plain, ((e16 = op1(e14, e12)) | ~ spl42_133), inference(avatar_component_clause, [], [f2137])).
fof(f276, plain, (e12 = op1(op1(op1(e14, e12), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f19480, plain, ((e15 = op1(op1(e16, e14), e14)) | ~ spl42_112), inference(backward_demodulation, [], [f297, f2052])).
fof(f297, plain, (e15 = op1(op1(op1(e14, e15), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f20457, plain, (~ spl42_139 | ~ spl42_146), inference(avatar_split_clause, [], [f20452, f2191, f2162])).
fof(f2191, plain, (spl42_146 <=> (e15 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_146])])).
fof(f20452, plain, (~ (e15 = op1(e14, e11)) | ~ spl42_146), inference(backward_demodulation, [], [f734, f2193])).
fof(f2193, plain, ((e15 = op1(e14, e10)) | ~ spl42_146), inference(avatar_component_clause, [], [f2191])).
fof(f734, plain, ~ (op1(e14, e10) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f20449, plain, (~ spl42_39 | spl42_149 | ~ spl42_151), inference(avatar_contradiction_clause, [], [f20448])).
fof(f20448, plain, ($false | (~ spl42_39 | spl42_149 | ~ spl42_151)), inference(subsumption_resolution, [], [f20442, f2205])).
fof(f2205, plain, (~ (e11 = op1(e13, e16)) | spl42_149), inference(avatar_component_clause, [], [f2204])).
fof(f20442, plain, ((e11 = op1(e13, e16)) | (~ spl42_39 | ~ spl42_151)), inference(backward_demodulation, [], [f12233, f2214])).
fof(f2214, plain, ((e13 = op1(e13, e16)) | ~ spl42_151), inference(avatar_component_clause, [], [f2212])).
fof(f2212, plain, (spl42_151 <=> (e13 = op1(e13, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_151])])).
fof(f20402, plain, (spl42_121 | ~ spl42_26 | ~ spl42_189), inference(avatar_split_clause, [], [f20401, f2369, f1694, f2088])).
fof(f2088, plain, (spl42_121 <=> (e11 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_121])])).
fof(f2369, plain, (spl42_189 <=> (e16 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_189])])).
fof(f20401, plain, ((e11 = op1(e14, e13)) | (~ spl42_26 | ~ spl42_189)), inference(forward_demodulation, [], [f20394, f1696])).
fof(f20394, plain, ((e11 = op1(op1(e16, e13), e13)) | ~ spl42_189), inference(backward_demodulation, [], [f268, f2371])).
fof(f2371, plain, ((e16 = op1(e13, e11)) | ~ spl42_189), inference(avatar_component_clause, [], [f2369])).
fof(f268, plain, (e11 = op1(op1(op1(e13, e11), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f20360, plain, (~ spl42_74 | ~ spl42_221), inference(avatar_split_clause, [], [f20352, f2502, f1893])).
fof(f2502, plain, (spl42_221 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_221])])).
fof(f20352, plain, (~ (e13 = op1(e15, e13)) | ~ spl42_221), inference(backward_demodulation, [], [f579, f2504])).
fof(f2504, plain, ((e13 = op1(e12, e13)) | ~ spl42_221), inference(avatar_component_clause, [], [f2502])).
fof(f20314, plain, (spl42_88 | ~ spl42_267 | ~ spl42_335), inference(avatar_contradiction_clause, [], [f20313])).
fof(f20313, plain, ($false | (spl42_88 | ~ spl42_267 | ~ spl42_335)), inference(subsumption_resolution, [], [f20312, f1952])).
fof(f1952, plain, (~ (e13 = op1(e15, e11)) | spl42_88), inference(avatar_component_clause, [], [f1951])).
fof(f1951, plain, (spl42_88 <=> (e13 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_88])])).
fof(f20312, plain, ((e13 = op1(e15, e11)) | (~ spl42_267 | ~ spl42_335)), inference(forward_demodulation, [], [f20301, f2976])).
fof(f2976, plain, ((e15 = op1(e10, e11)) | ~ spl42_335), inference(avatar_component_clause, [], [f2974])).
fof(f2974, plain, (spl42_335 <=> (e15 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_335])])).
fof(f20301, plain, ((e13 = op1(op1(e10, e11), e11)) | ~ spl42_267), inference(backward_demodulation, [], [f280, f2695])).
fof(f20282, plain, (~ spl42_184 | ~ spl42_282), inference(avatar_split_clause, [], [f20269, f2755, f2349])).
fof(f2755, plain, (spl42_282 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_282])])).
fof(f20269, plain, (~ (e11 = op1(e13, e11)) | ~ spl42_282), inference(backward_demodulation, [], [f531, f2757])).
fof(f2757, plain, ((e11 = op1(e11, e11)) | ~ spl42_282), inference(avatar_component_clause, [], [f2755])).
fof(f531, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f20263, plain, (~ spl42_8 | spl42_202 | ~ spl42_297), inference(avatar_contradiction_clause, [], [f20262])).
fof(f20262, plain, ($false | (~ spl42_8 | spl42_202 | ~ spl42_297)), inference(subsumption_resolution, [], [f20256, f2424])).
fof(f2424, plain, (~ (e15 = op1(e12, e16)) | spl42_202), inference(avatar_component_clause, [], [f2423])).
fof(f2423, plain, (spl42_202 <=> (e15 = op1(e12, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_202])])).
fof(f20256, plain, ((e15 = op1(e12, e16)) | (~ spl42_8 | ~ spl42_297)), inference(backward_demodulation, [], [f14670, f2819])).
fof(f2819, plain, ((e12 = op1(e10, e16)) | ~ spl42_297), inference(avatar_component_clause, [], [f2817])).
fof(f2817, plain, (spl42_297 <=> (e12 = op1(e10, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_297])])).
fof(f14670, plain, ((e15 = op1(op1(e10, e16), e16)) | ~ spl42_8), inference(backward_demodulation, [], [f299, f1622])).
fof(f299, plain, (e15 = op1(op1(op1(e16, e15), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f20188, plain, (~ spl42_139 | ~ spl42_335), inference(avatar_split_clause, [], [f20179, f2974, f2162])).
fof(f20179, plain, (~ (e15 = op1(e14, e11)) | ~ spl42_335), inference(backward_demodulation, [], [f527, f2976])).
fof(f527, plain, ~ (op1(e10, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f20166, plain, (~ spl42_321 | ~ spl42_342), inference(avatar_split_clause, [], [f20155, f3003, f2916])).
fof(f3003, plain, (spl42_342 <=> (op1(e10, e10) = e15)), introduced(avatar_definition, [new_symbols(naming, [spl42_342])])).
fof(f20155, plain, (~ (e15 = op1(e10, e13)) | ~ spl42_342), inference(backward_demodulation, [], [f652, f3005])).
fof(f3005, plain, ((op1(e10, e10) = e15) | ~ spl42_342), inference(avatar_component_clause, [], [f3003])).
fof(f652, plain, ~ (op1(e10, e10) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f20144, plain, (~ spl42_359 | ~ spl42_493 | ~ spl42_592), inference(avatar_contradiction_clause, [], [f20143])).
fof(f20143, plain, ($false | (~ spl42_359 | ~ spl42_493 | ~ spl42_592)), inference(subsumption_resolution, [], [f20142, f1124])).
fof(f1124, plain, ~ (e22 = e24), inference(cnf_transformation, [], [f10])).
fof(f20142, plain, ((e22 = e24) | (~ spl42_359 | ~ spl42_493 | ~ spl42_592)), inference(forward_demodulation, [], [f20141, f3729])).
fof(f3729, plain, ((e22 = op2(e23, e26)) | ~ spl42_493), inference(avatar_component_clause, [], [f3727])).
fof(f3727, plain, (spl42_493 <=> (e22 = op2(e23, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_493])])).
fof(f20141, plain, ((e24 = op2(e23, e26)) | (~ spl42_359 | ~ spl42_592)), inference(forward_demodulation, [], [f20133, f4139])).
fof(f20133, plain, ((e24 = op2(op2(e21, e26), e26)) | ~ spl42_359), inference(backward_demodulation, [], [f488, f3174])).
fof(f488, plain, (e24 = op2(op2(op2(e26, e24), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f20132, plain, (~ spl42_369 | ~ spl42_371), inference(avatar_contradiction_clause, [], [f20131])).
fof(f20131, plain, ($false | (~ spl42_369 | ~ spl42_371)), inference(subsumption_resolution, [], [f20130, f1131])).
fof(f20130, plain, ((e24 = e26) | (~ spl42_369 | ~ spl42_371)), inference(forward_demodulation, [], [f3223, f3215])).
fof(f3223, plain, ((e26 = op2(e26, e23)) | ~ spl42_371), inference(avatar_component_clause, [], [f3221])).
fof(f3221, plain, (spl42_371 <=> (e26 = op2(e26, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_371])])).
fof(f20112, plain, (~ spl42_393 | ~ spl42_502 | ~ spl42_648), inference(avatar_contradiction_clause, [], [f20111])).
fof(f20111, plain, ($false | (~ spl42_393 | ~ spl42_502 | ~ spl42_648)), inference(subsumption_resolution, [], [f20110, f1131])).
fof(f20110, plain, ((e24 = e26) | (~ spl42_393 | ~ spl42_502 | ~ spl42_648)), inference(forward_demodulation, [], [f20109, f3766])).
fof(f3766, plain, ((e24 = op2(e23, e25)) | ~ spl42_502), inference(avatar_component_clause, [], [f3764])).
fof(f3764, plain, (spl42_502 <=> (e24 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_502])])).
fof(f20109, plain, ((e26 = op2(e23, e25)) | (~ spl42_393 | ~ spl42_648)), inference(forward_demodulation, [], [f20100, f4371])).
fof(f4371, plain, ((e23 = op2(e20, e25)) | ~ spl42_648), inference(avatar_component_clause, [], [f4369])).
fof(f4369, plain, (spl42_648 <=> (e23 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_648])])).
fof(f20081, plain, (spl42_400 | ~ spl42_437 | ~ spl42_552), inference(avatar_contradiction_clause, [], [f20080])).
fof(f20080, plain, ($false | (spl42_400 | ~ spl42_437 | ~ spl42_552)), inference(subsumption_resolution, [], [f20079, f3343])).
fof(f3343, plain, (~ (e20 = op2(e25, e25)) | spl42_400), inference(avatar_component_clause, [], [f3342])).
fof(f3342, plain, (spl42_400 <=> (e20 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_400])])).
fof(f20079, plain, ((e20 = op2(e25, e25)) | (~ spl42_437 | ~ spl42_552)), inference(forward_demodulation, [], [f20070, f3973])).
fof(f20070, plain, ((e20 = op2(op2(e22, e25), e25)) | ~ spl42_437), inference(backward_demodulation, [], [f459, f3497])).
fof(f3497, plain, ((e22 = op2(e25, e20)) | ~ spl42_437), inference(avatar_component_clause, [], [f3495])).
fof(f3495, plain, (spl42_437 <=> (e22 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_437])])).
fof(f20069, plain, (~ spl42_369 | spl42_445 | ~ spl42_446), inference(avatar_contradiction_clause, [], [f20068])).
fof(f20068, plain, ($false | (~ spl42_369 | spl42_445 | ~ spl42_446)), inference(subsumption_resolution, [], [f20062, f3529])).
fof(f3529, plain, (~ (e23 = op2(e24, e26)) | spl42_445), inference(avatar_component_clause, [], [f3528])).
fof(f3528, plain, (spl42_445 <=> (e23 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_445])])).
fof(f20062, plain, ((e23 = op2(e24, e26)) | (~ spl42_369 | ~ spl42_446)), inference(backward_demodulation, [], [f12623, f3534])).
fof(f3534, plain, ((e24 = op2(e24, e26)) | ~ spl42_446), inference(avatar_component_clause, [], [f3532])).
fof(f3532, plain, (spl42_446 <=> (e24 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_446])])).
fof(f12623, plain, ((e23 = op2(op2(e24, e26), e26)) | ~ spl42_369), inference(backward_demodulation, [], [f481, f3215])).
fof(f481, plain, (e23 = op2(op2(op2(e26, e23), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f19951, plain, (~ spl42_387 | spl42_491 | ~ spl42_592), inference(avatar_contradiction_clause, [], [f19950])).
fof(f19950, plain, ($false | (~ spl42_387 | spl42_491 | ~ spl42_592)), inference(subsumption_resolution, [], [f19944, f3720])).
fof(f3720, plain, (~ (e20 = op2(e23, e26)) | spl42_491), inference(avatar_component_clause, [], [f3719])).
fof(f19944, plain, ((e20 = op2(e23, e26)) | (~ spl42_387 | ~ spl42_592)), inference(backward_demodulation, [], [f19238, f4139])).
fof(f19238, plain, ((e20 = op2(op2(e21, e26), e26)) | ~ spl42_387), inference(backward_demodulation, [], [f460, f3290])).
fof(f3290, plain, ((e21 = op2(e26, e20)) | ~ spl42_387), inference(avatar_component_clause, [], [f3288])).
fof(f19948, plain, (~ spl42_382 | ~ spl42_532 | ~ spl42_592), inference(avatar_contradiction_clause, [], [f19947])).
fof(f19947, plain, ($false | (~ spl42_382 | ~ spl42_532 | ~ spl42_592)), inference(subsumption_resolution, [], [f19946, f1129])).
fof(f19946, plain, ((e23 = e26) | (~ spl42_382 | ~ spl42_532 | ~ spl42_592)), inference(forward_demodulation, [], [f19945, f3269])).
fof(f19945, plain, ((e26 = op2(e26, e21)) | (~ spl42_532 | ~ spl42_592)), inference(forward_demodulation, [], [f19937, f3890])).
fof(f3890, plain, ((e26 = op2(e23, e21)) | ~ spl42_532), inference(avatar_component_clause, [], [f3888])).
fof(f3888, plain, (spl42_532 <=> (e26 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_532])])).
fof(f19937, plain, ((e26 = op2(op2(e23, e21), e21)) | ~ spl42_592), inference(backward_demodulation, [], [f497, f4139])).
fof(f19873, plain, (~ spl42_661 | ~ spl42_664), inference(avatar_contradiction_clause, [], [f19872])).
fof(f19872, plain, ($false | (~ spl42_661 | ~ spl42_664)), inference(subsumption_resolution, [], [f19871, f1125])).
fof(f19871, plain, ((e22 = e25) | (~ spl42_661 | ~ spl42_664)), inference(forward_demodulation, [], [f4437, f4425])).
fof(f4425, plain, ((e22 = op2(e20, e23)) | ~ spl42_661), inference(avatar_component_clause, [], [f4423])).
fof(f4423, plain, (spl42_661 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_661])])).
fof(f19718, plain, (spl42_678 | ~ spl42_575 | ~ spl42_598), inference(avatar_split_clause, [], [f19015, f4162, f4067, f4493])).
fof(f4067, plain, (spl42_575 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_575])])).
fof(f19015, plain, ((e25 = op2(e20, e21)) | (~ spl42_575 | ~ spl42_598)), inference(backward_demodulation, [], [f14033, f4069])).
fof(f4069, plain, ((e20 = op2(e22, e21)) | ~ spl42_575), inference(avatar_component_clause, [], [f4067])).
fof(f19693, plain, (~ spl42_351 | spl42_594 | ~ spl42_639), inference(avatar_contradiction_clause, [], [f19692])).
fof(f19692, plain, ($false | (~ spl42_351 | spl42_594 | ~ spl42_639)), inference(subsumption_resolution, [], [f19691, f4146])).
fof(f4146, plain, (~ (e25 = op2(e21, e26)) | spl42_594), inference(avatar_component_clause, [], [f4145])).
fof(f4145, plain, (spl42_594 <=> (e25 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_594])])).
fof(f19691, plain, ((e25 = op2(e21, e26)) | (~ spl42_351 | ~ spl42_639)), inference(forward_demodulation, [], [f12583, f4334])).
fof(f4334, plain, ((e21 = op2(e20, e26)) | ~ spl42_639), inference(avatar_component_clause, [], [f4332])).
fof(f12583, plain, ((e25 = op2(op2(e20, e26), e26)) | ~ spl42_351), inference(forward_demodulation, [], [f495, f3141])).
fof(f495, plain, (e25 = op2(op2(op2(e26, e25), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f19643, plain, (spl42_528 | ~ spl42_382 | ~ spl42_623), inference(avatar_split_clause, [], [f18947, f4265, f3267, f3872])).
fof(f3872, plain, (spl42_528 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_528])])).
fof(f4265, plain, (spl42_623 <=> (e26 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_623])])).
fof(f18947, plain, ((e22 = op2(e23, e21)) | (~ spl42_382 | ~ spl42_623)), inference(forward_demodulation, [], [f18942, f3269])).
fof(f18942, plain, ((e22 = op2(op2(e26, e21), e21)) | ~ spl42_623), inference(backward_demodulation, [], [f469, f4267])).
fof(f4267, plain, ((e26 = op2(e21, e22)) | ~ spl42_623), inference(avatar_component_clause, [], [f4265])).
fof(f469, plain, (e22 = op2(op2(op2(e21, e22), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f19620, plain, (~ spl42_373 | ~ spl42_387), inference(avatar_split_clause, [], [f19619, f3288, f3230])).
fof(f19619, plain, (~ (e21 = op2(e26, e22)) | ~ spl42_387), inference(forward_demodulation, [], [f1071, f3290])).
fof(f19604, plain, (~ spl42_303 | ~ spl42_58), inference(avatar_split_clause, [], [f13778, f1827, f2842])).
fof(f2842, plain, (spl42_303 <=> (e11 = op1(e10, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_303])])).
fof(f13778, plain, (~ (e11 = op1(e10, e15)) | ~ spl42_58), inference(forward_demodulation, [], [f612, f1829])).
fof(f612, plain, ~ (op1(e10, e15) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f19595, plain, (~ spl42_249 | ~ spl42_291), inference(avatar_split_clause, [], [f19594, f2792, f2618])).
fof(f2792, plain, (spl42_291 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_291])])).
fof(f19594, plain, (~ (e13 = op1(e11, e16)) | ~ spl42_291), inference(forward_demodulation, [], [f676, f2794])).
fof(f2794, plain, ((e13 = op1(e11, e10)) | ~ spl42_291), inference(avatar_component_clause, [], [f2792])).
fof(f676, plain, ~ (op1(e11, e10) = op1(e11, e16)), inference(cnf_transformation, [], [f7])).
fof(f19579, plain, (~ spl42_160 | ~ spl42_209), inference(avatar_split_clause, [], [f14445, f2452, f2249])).
fof(f2249, plain, (spl42_160 <=> (e15 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_160])])).
fof(f14445, plain, (~ (e15 = op1(e13, e15)) | ~ spl42_209), inference(backward_demodulation, [], [f619, f2454])).
fof(f619, plain, ~ (op1(e12, e15) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f19561, plain, (~ spl42_74 | ~ spl42_123), inference(avatar_split_clause, [], [f17415, f2096, f1893])).
fof(f2096, plain, (spl42_123 <=> (e13 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_123])])).
fof(f17415, plain, (~ (e13 = op1(e15, e13)) | ~ spl42_123), inference(backward_demodulation, [], [f584, f2098])).
fof(f2098, plain, ((e13 = op1(e14, e13)) | ~ spl42_123), inference(avatar_component_clause, [], [f2096])).
fof(f584, plain, ~ (op1(e14, e13) = op1(e15, e13)), inference(cnf_transformation, [], [f7])).
fof(f19385, plain, (~ spl42_226 | ~ spl42_228), inference(avatar_contradiction_clause, [], [f19384])).
fof(f19384, plain, ($false | (~ spl42_226 | ~ spl42_228)), inference(subsumption_resolution, [], [f19383, f1098])).
fof(f1098, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f9])).
fof(f19383, plain, ((e11 = e13) | (~ spl42_226 | ~ spl42_228)), inference(backward_demodulation, [], [f2533, f2525])).
fof(f2525, plain, ((e11 = op1(e12, e12)) | ~ spl42_226), inference(avatar_component_clause, [], [f2523])).
fof(f2523, plain, (spl42_226 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_226])])).
fof(f19363, plain, (~ spl42_249 | ~ spl42_270), inference(avatar_contradiction_clause, [], [f19362])).
fof(f19362, plain, ($false | (~ spl42_249 | ~ spl42_270)), inference(subsumption_resolution, [], [f19361, f1108])).
fof(f1108, plain, ~ (e13 = e16), inference(cnf_transformation, [], [f9])).
fof(f19361, plain, ((e13 = e16) | (~ spl42_249 | ~ spl42_270)), inference(forward_demodulation, [], [f19357, f19345])).
fof(f19345, plain, ((e13 = op1(op1(e13, e11), e11)) | ~ spl42_270), inference(backward_demodulation, [], [f280, f2707])).
fof(f2707, plain, ((e13 = op1(e11, e13)) | ~ spl42_270), inference(avatar_component_clause, [], [f2705])).
fof(f2705, plain, (spl42_270 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_270])])).
fof(f19357, plain, ((e16 = op1(op1(e13, e11), e11)) | ~ spl42_249), inference(backward_demodulation, [], [f301, f2620])).
fof(f301, plain, (e16 = op1(op1(op1(e11, e16), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f19356, plain, (~ spl42_254 | ~ spl42_255), inference(avatar_contradiction_clause, [], [f19355])).
fof(f19355, plain, ($false | (~ spl42_254 | ~ spl42_255)), inference(subsumption_resolution, [], [f19354, f1097])).
fof(f19354, plain, ((e11 = e12) | (~ spl42_254 | ~ spl42_255)), inference(backward_demodulation, [], [f2645, f2641])).
fof(f2641, plain, ((e11 = op1(e11, e15)) | ~ spl42_254), inference(avatar_component_clause, [], [f2639])).
fof(f2639, plain, (spl42_254 <=> (e11 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_254])])).
fof(f19325, plain, (~ spl42_8 | ~ spl42_295), inference(avatar_contradiction_clause, [], [f19324])).
fof(f19324, plain, ($false | (~ spl42_8 | ~ spl42_295)), inference(subsumption_resolution, [], [f19323, f1095])).
fof(f1095, plain, ~ (e10 = e15), inference(cnf_transformation, [], [f9])).
fof(f19323, plain, ((e10 = e15) | (~ spl42_8 | ~ spl42_295)), inference(forward_demodulation, [], [f19320, f2811])).
fof(f2811, plain, ((e10 = op1(e10, e16)) | ~ spl42_295), inference(avatar_component_clause, [], [f2809])).
fof(f2809, plain, (spl42_295 <=> (e10 = op1(e10, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_295])])).
fof(f19320, plain, ((e15 = op1(e10, e16)) | (~ spl42_8 | ~ spl42_295)), inference(backward_demodulation, [], [f14670, f2811])).
fof(f19258, plain, (spl42_363 | ~ spl42_364 | ~ spl42_455), inference(avatar_contradiction_clause, [], [f19257])).
fof(f19257, plain, ($false | (spl42_363 | ~ spl42_364 | ~ spl42_455)), inference(subsumption_resolution, [], [f19255, f3189])).
fof(f3189, plain, (~ (e25 = op2(e26, e24)) | spl42_363), inference(avatar_component_clause, [], [f3188])).
fof(f3188, plain, (spl42_363 <=> (e25 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_363])])).
fof(f19255, plain, ((e25 = op2(e26, e24)) | (~ spl42_364 | ~ spl42_455)), inference(backward_demodulation, [], [f19164, f3194])).
fof(f3194, plain, ((e26 = op2(e26, e24)) | ~ spl42_364), inference(avatar_component_clause, [], [f3192])).
fof(f19164, plain, ((e25 = op2(op2(e26, e24), e24)) | ~ spl42_455), inference(backward_demodulation, [], [f493, f3571])).
fof(f19220, plain, (~ spl42_401 | ~ spl42_412 | ~ spl42_598), inference(avatar_contradiction_clause, [], [f19219])).
fof(f19219, plain, ($false | (~ spl42_401 | ~ spl42_412 | ~ spl42_598)), inference(subsumption_resolution, [], [f19218, f1124])).
fof(f19218, plain, ((e22 = e24) | (~ spl42_401 | ~ spl42_412 | ~ spl42_598)), inference(forward_demodulation, [], [f19217, f4164])).
fof(f19217, plain, ((e24 = op2(e21, e25)) | (~ spl42_401 | ~ spl42_412)), inference(forward_demodulation, [], [f19211, f3348])).
fof(f19211, plain, ((e24 = op2(op2(e25, e25), e25)) | ~ spl42_412), inference(backward_demodulation, [], [f487, f3393])).
fof(f3393, plain, ((e25 = op2(e25, e24)) | ~ spl42_412), inference(avatar_component_clause, [], [f3391])).
fof(f19190, plain, (~ spl42_435 | spl42_498 | ~ spl42_648), inference(avatar_contradiction_clause, [], [f19189])).
fof(f19189, plain, ($false | (~ spl42_435 | spl42_498 | ~ spl42_648)), inference(subsumption_resolution, [], [f19188, f3749])).
fof(f3749, plain, (~ (e20 = op2(e23, e25)) | spl42_498), inference(avatar_component_clause, [], [f3748])).
fof(f3748, plain, (spl42_498 <=> (e20 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_498])])).
fof(f19188, plain, ((e20 = op2(e23, e25)) | (~ spl42_435 | ~ spl42_648)), inference(forward_demodulation, [], [f19174, f4371])).
fof(f19174, plain, ((e20 = op2(op2(e20, e25), e25)) | ~ spl42_435), inference(backward_demodulation, [], [f459, f3489])).
fof(f3489, plain, ((e20 = op2(e25, e20)) | ~ spl42_435), inference(avatar_component_clause, [], [f3487])).
fof(f3487, plain, (spl42_435 <=> (e20 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_435])])).
fof(f19173, plain, (spl42_347 | ~ spl42_369 | ~ spl42_448), inference(avatar_split_clause, [], [f19169, f3540, f3213, f3122])).
fof(f3122, plain, (spl42_347 <=> (e23 = op2(e26, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_347])])).
fof(f3540, plain, (spl42_448 <=> (e26 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_448])])).
fof(f19169, plain, ((e23 = op2(e26, e26)) | (~ spl42_369 | ~ spl42_448)), inference(backward_demodulation, [], [f12623, f3542])).
fof(f3542, plain, ((e26 = op2(e24, e26)) | ~ spl42_448), inference(avatar_component_clause, [], [f3540])).
fof(f19144, plain, (spl42_492 | ~ spl42_494 | ~ spl42_698), inference(avatar_contradiction_clause, [], [f19143])).
fof(f19143, plain, ($false | (spl42_492 | ~ spl42_494 | ~ spl42_698)), inference(subsumption_resolution, [], [f19138, f3724])).
fof(f3724, plain, (~ (e21 = op2(e23, e26)) | spl42_492), inference(avatar_component_clause, [], [f3723])).
fof(f19138, plain, ((e21 = op2(e23, e26)) | (~ spl42_494 | ~ spl42_698)), inference(backward_demodulation, [], [f15201, f3733])).
fof(f3733, plain, ((e23 = op2(e23, e26)) | ~ spl42_494), inference(avatar_component_clause, [], [f3731])).
fof(f3731, plain, (spl42_494 <=> (e23 = op2(e23, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_494])])).
fof(f18984, plain, (~ spl42_591 | ~ spl42_598), inference(avatar_contradiction_clause, [], [f18983])).
fof(f18983, plain, ($false | (~ spl42_591 | ~ spl42_598)), inference(subsumption_resolution, [], [f18982, f1132])).
fof(f1132, plain, ~ (e25 = e26), inference(cnf_transformation, [], [f10])).
fof(f18982, plain, ((e25 = e26) | (~ spl42_591 | ~ spl42_598)), inference(forward_demodulation, [], [f18973, f14033])).
fof(f18973, plain, ((e26 = op2(op2(e22, e21), e21)) | ~ spl42_591), inference(backward_demodulation, [], [f497, f4135])).
fof(f4135, plain, ((e22 = op2(e21, e26)) | ~ spl42_591), inference(avatar_component_clause, [], [f4133])).
fof(f4133, plain, (spl42_591 <=> (e22 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_591])])).
fof(f18934, plain, (~ spl42_648 | ~ spl42_649), inference(avatar_contradiction_clause, [], [f18933])).
fof(f18933, plain, ($false | (~ spl42_648 | ~ spl42_649)), inference(subsumption_resolution, [], [f18932, f1127])).
fof(f1127, plain, ~ (e23 = e24), inference(cnf_transformation, [], [f10])).
fof(f18932, plain, ((e23 = e24) | (~ spl42_648 | ~ spl42_649)), inference(backward_demodulation, [], [f4375, f4371])).
fof(f18879, plain, (spl42_908 | ~ spl42_682), inference(avatar_split_clause, [], [f18878, f4510, f5942])).
fof(f18878, plain, ((e22 = h1(e15)) | ~ spl42_682), inference(backward_demodulation, [], [f1204, f4512])).
fof(f18875, plain, (spl42_735 | ~ spl42_401 | ~ spl42_733 | ~ spl42_852), inference(avatar_split_clause, [], [f18874, f5661, f5063, f3346, f5073])).
fof(f5661, plain, (spl42_852 <=> (e25 = h2(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_852])])).
fof(f18874, plain, ((e21 = h6(e13)) | (~ spl42_401 | ~ spl42_733 | ~ spl42_852)), inference(forward_demodulation, [], [f18859, f3348])).
fof(f18859, plain, ((op2(e25, e25) = h6(e13)) | (~ spl42_733 | ~ spl42_852)), inference(backward_demodulation, [], [f14049, f5662])).
fof(f5662, plain, ((e25 = h2(e15)) | ~ spl42_852), inference(avatar_component_clause, [], [f5661])).
fof(f18839, plain, (~ spl42_677 | ~ spl42_649), inference(avatar_split_clause, [], [f18838, f4373, f4489])).
fof(f18838, plain, (~ (e24 = op2(e20, e21)) | ~ spl42_649), inference(forward_demodulation, [], [f953, f4375])).
fof(f953, plain, ~ (op2(e20, e21) = op2(e20, e25)), inference(cnf_transformation, [], [f8])).
fof(f18837, plain, (~ spl42_676 | ~ spl42_698), inference(avatar_split_clause, [], [f18836, f4872, f4485])).
fof(f4485, plain, (spl42_676 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_676])])).
fof(f18836, plain, (~ (e23 = op2(e20, e21)) | ~ spl42_698), inference(forward_demodulation, [], [f4791, f4873])).
fof(f4791, plain, ~ (op2(e20, e21) = h7(e13)), inference(backward_demodulation, [], [f823, f4789])).
fof(f823, plain, ~ (op2(e20, e21) = op2(e26, e21)), inference(cnf_transformation, [], [f8])).
fof(f18817, plain, (~ spl42_351 | ~ spl42_638), inference(avatar_contradiction_clause, [], [f18816])).
fof(f18816, plain, ($false | (~ spl42_351 | ~ spl42_638)), inference(subsumption_resolution, [], [f18815, f1116])).
fof(f1116, plain, ~ (e20 = e25), inference(cnf_transformation, [], [f10])).
fof(f18815, plain, ((e20 = e25) | (~ spl42_351 | ~ spl42_638)), inference(forward_demodulation, [], [f18814, f4330])).
fof(f4330, plain, ((e20 = op2(e20, e26)) | ~ spl42_638), inference(avatar_component_clause, [], [f4328])).
fof(f4328, plain, (spl42_638 <=> (e20 = op2(e20, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_638])])).
fof(f18814, plain, ((e25 = op2(e20, e26)) | (~ spl42_351 | ~ spl42_638)), inference(forward_demodulation, [], [f12583, f4330])).
fof(f18787, plain, (spl42_607 | ~ spl42_481), inference(avatar_split_clause, [], [f17973, f3677, f4199])).
fof(f3677, plain, (spl42_481 <=> (e24 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_481])])).
fof(f17973, plain, ((e24 = op2(e21, e24)) | ~ spl42_481), inference(backward_demodulation, [], [f4728, f17972])).
fof(f17972, plain, ((e21 = op2(h5(e15), e24)) | ~ spl42_481), inference(forward_demodulation, [], [f17961, f1232])).
fof(f17961, plain, ((e21 = op2(op2(e24, e24), e24)) | ~ spl42_481), inference(backward_demodulation, [], [f465, f3679])).
fof(f3679, plain, ((e24 = op2(e24, e21)) | ~ spl42_481), inference(avatar_component_clause, [], [f3677])).
fof(f4728, plain, (e24 = op2(op2(h5(e15), e24), e24)), inference(backward_demodulation, [], [f486, f1232])).
fof(f486, plain, (e24 = op2(op2(op2(e24, e24), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f18686, plain, (~ spl42_502 | ~ spl42_649), inference(avatar_split_clause, [], [f18685, f4373, f3764])).
fof(f18685, plain, (~ (e24 = op2(e23, e25)) | ~ spl42_649), inference(forward_demodulation, [], [f904, f4375])).
fof(f904, plain, ~ (op2(e20, e25) = op2(e23, e25)), inference(cnf_transformation, [], [f8])).
fof(f18684, plain, (~ spl42_499 | ~ spl42_733), inference(avatar_split_clause, [], [f15978, f5063, f3752])).
fof(f3752, plain, (spl42_499 <=> (e21 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_499])])).
fof(f15978, plain, (~ (e21 = op2(e23, e25)) | ~ spl42_733), inference(forward_demodulation, [], [f4753, f5064])).
fof(f4753, plain, ~ (op2(e23, e25) = h6(e15)), inference(backward_demodulation, [], [f918, f1239])).
fof(f918, plain, ~ (op2(e23, e25) = op2(e25, e25)), inference(cnf_transformation, [], [f8])).
fof(f18646, plain, (~ spl42_333 | ~ spl42_39), inference(avatar_split_clause, [], [f18645, f1748, f2966])).
fof(f2966, plain, (spl42_333 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_333])])).
fof(f18645, plain, (~ (e13 = op1(e10, e11)) | ~ spl42_39), inference(forward_demodulation, [], [f529, f1750])).
fof(f529, plain, ~ (op1(e10, e11) = op1(e16, e11)), inference(cnf_transformation, [], [f7])).
fof(f18631, plain, (spl42_315 | ~ spl42_100 | ~ spl42_260), inference(avatar_split_clause, [], [f18411, f2664, f2001, f2891])).
fof(f2891, plain, (spl42_315 <=> (e16 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_315])])).
fof(f2664, plain, (spl42_260 <=> (e10 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_260])])).
fof(f18411, plain, ((e16 = op1(e10, e14)) | (~ spl42_100 | ~ spl42_260)), inference(forward_demodulation, [], [f18403, f2666])).
fof(f2666, plain, ((e10 = op1(e11, e14)) | ~ spl42_260), inference(avatar_component_clause, [], [f2664])).
fof(f18403, plain, ((e16 = op1(op1(e11, e14), e14)) | ~ spl42_100), inference(backward_demodulation, [], [f304, f2003])).
fof(f18622, plain, (spl42_335 | ~ spl42_232 | ~ spl42_255), inference(avatar_split_clause, [], [f18621, f2643, f2548, f2974])).
fof(f18621, plain, ((e15 = op1(e10, e11)) | (~ spl42_232 | ~ spl42_255)), inference(forward_demodulation, [], [f13742, f2550])).
fof(f2550, plain, ((e10 = op1(e12, e11)) | ~ spl42_232), inference(avatar_component_clause, [], [f2548])).
fof(f18599, plain, (spl42_249 | ~ spl42_26 | ~ spl42_100), inference(avatar_split_clause, [], [f18410, f2001, f1694, f2618])).
fof(f18410, plain, ((e13 = op1(e11, e16)) | (~ spl42_26 | ~ spl42_100)), inference(backward_demodulation, [], [f12209, f2003])).
fof(f18583, plain, (~ spl42_52 | ~ spl42_58 | ~ spl42_209), inference(avatar_contradiction_clause, [], [f18582])).
fof(f18582, plain, ($false | (~ spl42_52 | ~ spl42_58 | ~ spl42_209)), inference(subsumption_resolution, [], [f18581, f1101])).
fof(f1101, plain, ~ (e11 = e16), inference(cnf_transformation, [], [f9])).
fof(f18581, plain, ((e11 = e16) | (~ spl42_52 | ~ spl42_58 | ~ spl42_209)), inference(forward_demodulation, [], [f18580, f1829])).
fof(f18580, plain, ((e16 = op1(e15, e15)) | (~ spl42_52 | ~ spl42_209)), inference(backward_demodulation, [], [f18450, f2454])).
fof(f18450, plain, ((e16 = op1(op1(e12, e15), e15)) | ~ spl42_52), inference(backward_demodulation, [], [f305, f1804])).
fof(f1804, plain, ((e12 = op1(e15, e16)) | ~ spl42_52), inference(avatar_component_clause, [], [f1802])).
fof(f1802, plain, (spl42_52 <=> (e12 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_52])])).
fof(f18578, plain, (~ spl42_187 | ~ spl42_152), inference(avatar_split_clause, [], [f18577, f2216, f2361])).
fof(f2361, plain, (spl42_187 <=> (e14 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_187])])).
fof(f18577, plain, (~ (e14 = op1(e13, e11)) | ~ spl42_152), inference(forward_demodulation, [], [f723, f2218])).
fof(f723, plain, ~ (op1(e13, e11) = op1(e13, e16)), inference(cnf_transformation, [], [f7])).
fof(f18576, plain, (spl42_184 | ~ spl42_39 | ~ spl42_287), inference(avatar_split_clause, [], [f18183, f2775, f1748, f2349])).
fof(f18183, plain, ((e11 = op1(e13, e11)) | (~ spl42_39 | ~ spl42_287)), inference(forward_demodulation, [], [f18177, f1750])).
fof(f18177, plain, ((e11 = op1(op1(e16, e11), e11)) | ~ spl42_287), inference(backward_demodulation, [], [f266, f2777])).
fof(f18563, plain, (spl42_73 | ~ spl42_176 | ~ spl42_321), inference(avatar_contradiction_clause, [], [f18562])).
fof(f18562, plain, ($false | (spl42_73 | ~ spl42_176 | ~ spl42_321)), inference(subsumption_resolution, [], [f18561, f1890])).
fof(f1890, plain, (~ (e12 = op1(e15, e13)) | spl42_73), inference(avatar_component_clause, [], [f1889])).
fof(f1889, plain, (spl42_73 <=> (e12 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_73])])).
fof(f18561, plain, ((e12 = op1(e15, e13)) | (~ spl42_176 | ~ spl42_321)), inference(forward_demodulation, [], [f18560, f2918])).
fof(f18560, plain, ((e12 = op1(op1(e10, e13), e13)) | ~ spl42_176), inference(forward_demodulation, [], [f275, f2318])).
fof(f2318, plain, ((e10 = op1(e13, e12)) | ~ spl42_176), inference(avatar_component_clause, [], [f2316])).
fof(f2316, plain, (spl42_176 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_176])])).
fof(f18530, plain, (~ spl42_166 | ~ spl42_152), inference(avatar_split_clause, [], [f18529, f2216, f2274])).
fof(f2274, plain, (spl42_166 <=> (e14 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_166])])).
fof(f18529, plain, (~ (e14 = op1(e13, e14)) | ~ spl42_152), inference(forward_demodulation, [], [f732, f2218])).
fof(f732, plain, ~ (op1(e13, e14) = op1(e13, e16)), inference(cnf_transformation, [], [f7])).
fof(f18520, plain, (spl42_132 | ~ spl42_82 | ~ spl42_209), inference(avatar_split_clause, [], [f18436, f2452, f1926, f2133])).
fof(f2133, plain, (spl42_132 <=> (e15 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_132])])).
fof(f1926, plain, (spl42_82 <=> (e14 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_82])])).
fof(f18436, plain, ((e15 = op1(e14, e12)) | (~ spl42_82 | ~ spl42_209)), inference(backward_demodulation, [], [f14444, f1928])).
fof(f1928, plain, ((e14 = op1(e15, e12)) | ~ spl42_82), inference(avatar_component_clause, [], [f1926])).
fof(f14444, plain, ((e15 = op1(op1(e15, e12), e12)) | ~ spl42_209), inference(backward_demodulation, [], [f295, f2454])).
fof(f295, plain, (e15 = op1(op1(op1(e12, e15), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f18449, plain, (~ spl42_58 | ~ spl42_62), inference(avatar_contradiction_clause, [], [f18448])).
fof(f18448, plain, ($false | (~ spl42_58 | ~ spl42_62)), inference(subsumption_resolution, [], [f18447, f1100])).
fof(f1100, plain, ~ (e11 = e15), inference(cnf_transformation, [], [f9])).
fof(f18447, plain, ((e11 = e15) | (~ spl42_58 | ~ spl42_62)), inference(forward_demodulation, [], [f1845, f1829])).
fof(f1845, plain, ((e15 = op1(e15, e15)) | ~ spl42_62), inference(avatar_component_clause, [], [f1843])).
fof(f1843, plain, (spl42_62 <=> (e15 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_62])])).
fof(f18210, plain, (~ spl42_170 | ~ spl42_268), inference(avatar_contradiction_clause, [], [f18209])).
fof(f18209, plain, ($false | (~ spl42_170 | ~ spl42_268)), inference(subsumption_resolution, [], [f18208, f1098])).
fof(f18208, plain, ((e11 = e13) | (~ spl42_170 | ~ spl42_268)), inference(forward_demodulation, [], [f18206, f2699])).
fof(f18206, plain, ((e13 = op1(e11, e13)) | (~ spl42_170 | ~ spl42_268)), inference(backward_demodulation, [], [f16524, f2699])).
fof(f16524, plain, ((e13 = op1(op1(e11, e13), e13)) | ~ spl42_170), inference(backward_demodulation, [], [f282, f2293])).
fof(f2293, plain, ((e11 = op1(e13, e13)) | ~ spl42_170), inference(avatar_component_clause, [], [f2291])).
fof(f2291, plain, (spl42_170 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_170])])).
fof(f282, plain, (e13 = op1(op1(op1(e13, e13), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f18197, plain, (~ spl42_83 | ~ spl42_279), inference(avatar_split_clause, [], [f18192, f2742, f1930])).
fof(f2742, plain, (spl42_279 <=> (e15 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_279])])).
fof(f18192, plain, (~ (e15 = op1(e15, e12)) | ~ spl42_279), inference(backward_demodulation, [], [f554, f2744])).
fof(f2744, plain, ((e15 = op1(e11, e12)) | ~ spl42_279), inference(avatar_component_clause, [], [f2742])).
fof(f554, plain, ~ (op1(e11, e12) = op1(e15, e12)), inference(cnf_transformation, [], [f7])).
fof(f18159, plain, (~ spl42_8 | spl42_153 | ~ spl42_298), inference(avatar_contradiction_clause, [], [f18158])).
fof(f18158, plain, ($false | (~ spl42_8 | spl42_153 | ~ spl42_298)), inference(subsumption_resolution, [], [f18155, f2221])).
fof(f2221, plain, (~ (e15 = op1(e13, e16)) | spl42_153), inference(avatar_component_clause, [], [f2220])).
fof(f2220, plain, (spl42_153 <=> (e15 = op1(e13, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_153])])).
fof(f18155, plain, ((e15 = op1(e13, e16)) | (~ spl42_8 | ~ spl42_298)), inference(backward_demodulation, [], [f14670, f2823])).
fof(f2823, plain, ((e13 = op1(e10, e16)) | ~ spl42_298), inference(avatar_component_clause, [], [f2821])).
fof(f18143, plain, (~ spl42_159 | ~ spl42_306), inference(avatar_split_clause, [], [f18137, f2854, f2245])).
fof(f2245, plain, (spl42_159 <=> (e14 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_159])])).
fof(f18137, plain, (~ (e14 = op1(e13, e15)) | ~ spl42_306), inference(backward_demodulation, [], [f610, f2856])).
fof(f610, plain, ~ (op1(e10, e15) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f18124, plain, (~ spl42_347 | ~ spl42_349), inference(avatar_contradiction_clause, [], [f18123])).
fof(f18123, plain, ($false | (~ spl42_347 | ~ spl42_349)), inference(subsumption_resolution, [], [f18121, f1128])).
fof(f1128, plain, ~ (e23 = e25), inference(cnf_transformation, [], [f10])).
fof(f18121, plain, ((e23 = e25) | (~ spl42_347 | ~ spl42_349)), inference(backward_demodulation, [], [f3132, f3124])).
fof(f3124, plain, ((e23 = op2(e26, e26)) | ~ spl42_347), inference(avatar_component_clause, [], [f3122])).
fof(f18120, plain, (~ spl42_351 | ~ spl42_357), inference(avatar_contradiction_clause, [], [f18119])).
fof(f18119, plain, ($false | (~ spl42_351 | ~ spl42_357)), inference(subsumption_resolution, [], [f18118, f1117])).
fof(f18118, plain, ((e20 = e26) | (~ spl42_351 | ~ spl42_357)), inference(forward_demodulation, [], [f3165, f3141])).
fof(f3165, plain, ((e26 = op2(e26, e25)) | ~ spl42_357), inference(avatar_component_clause, [], [f3163])).
fof(f3163, plain, (spl42_357 <=> (e26 = op2(e26, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_357])])).
fof(f18100, plain, (~ spl42_401 | ~ spl42_405), inference(avatar_contradiction_clause, [], [f18099])).
fof(f18099, plain, ($false | (~ spl42_401 | ~ spl42_405)), inference(subsumption_resolution, [], [f18098, f1121])).
fof(f1121, plain, ~ (e21 = e25), inference(cnf_transformation, [], [f10])).
fof(f18098, plain, ((e21 = e25) | (~ spl42_401 | ~ spl42_405)), inference(forward_demodulation, [], [f3364, f3348])).
fof(f3364, plain, ((e25 = op2(e25, e25)) | ~ spl42_405), inference(avatar_component_clause, [], [f3362])).
fof(f3362, plain, (spl42_405 <=> (e25 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_405])])).
fof(f18097, plain, (spl42_361 | ~ spl42_413 | ~ spl42_468), inference(avatar_contradiction_clause, [], [f18096])).
fof(f18096, plain, ($false | (spl42_361 | ~ spl42_413 | ~ spl42_468)), inference(subsumption_resolution, [], [f18093, f3181])).
fof(f3181, plain, (~ (e23 = op2(e26, e24)) | spl42_361), inference(avatar_component_clause, [], [f3180])).
fof(f3180, plain, (spl42_361 <=> (e23 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_361])])).
fof(f18093, plain, ((e23 = op2(e26, e24)) | (~ spl42_413 | ~ spl42_468)), inference(backward_demodulation, [], [f18002, f3397])).
fof(f18002, plain, ((e23 = op2(op2(e25, e24), e24)) | ~ spl42_468), inference(backward_demodulation, [], [f479, f3625])).
fof(f3625, plain, ((e25 = op2(e24, e23)) | ~ spl42_468), inference(avatar_component_clause, [], [f3623])).
fof(f479, plain, (e23 = op2(op2(op2(e24, e23), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f18085, plain, (~ spl42_401 | ~ spl42_416 | ~ spl42_552), inference(avatar_contradiction_clause, [], [f18084])).
fof(f18084, plain, ($false | (~ spl42_401 | ~ spl42_416 | ~ spl42_552)), inference(subsumption_resolution, [], [f18083, f1119])).
fof(f1119, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f10])).
fof(f18083, plain, ((e21 = e23) | (~ spl42_401 | ~ spl42_416 | ~ spl42_552)), inference(forward_demodulation, [], [f18082, f3348])).
fof(f18082, plain, ((e23 = op2(e25, e25)) | (~ spl42_416 | ~ spl42_552)), inference(forward_demodulation, [], [f18073, f3973])).
fof(f18073, plain, ((e23 = op2(op2(e22, e25), e25)) | ~ spl42_416), inference(backward_demodulation, [], [f480, f3410])).
fof(f3410, plain, ((e22 = op2(e25, e23)) | ~ spl42_416), inference(avatar_component_clause, [], [f3408])).
fof(f3408, plain, (spl42_416 <=> (e22 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_416])])).
fof(f480, plain, (e23 = op2(op2(op2(e25, e23), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f18042, plain, (~ spl42_369 | ~ spl42_444 | ~ spl42_541), inference(avatar_contradiction_clause, [], [f18041])).
fof(f18041, plain, ($false | (~ spl42_369 | ~ spl42_444 | ~ spl42_541)), inference(subsumption_resolution, [], [f18040, f1119])).
fof(f18040, plain, ((e21 = e23) | (~ spl42_369 | ~ spl42_444 | ~ spl42_541)), inference(forward_demodulation, [], [f18035, f3928])).
fof(f3928, plain, ((e21 = op2(e22, e26)) | ~ spl42_541), inference(avatar_component_clause, [], [f3926])).
fof(f3926, plain, (spl42_541 <=> (e21 = op2(e22, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_541])])).
fof(f18035, plain, ((e23 = op2(e22, e26)) | (~ spl42_369 | ~ spl42_444)), inference(backward_demodulation, [], [f12623, f3526])).
fof(f3526, plain, ((e22 = op2(e24, e26)) | ~ spl42_444), inference(avatar_component_clause, [], [f3524])).
fof(f3524, plain, (spl42_444 <=> (e22 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_444])])).
fof(f17999, plain, (~ spl42_470 | spl42_556 | ~ spl42_654), inference(avatar_contradiction_clause, [], [f17998])).
fof(f17998, plain, ($false | (~ spl42_470 | spl42_556 | ~ spl42_654)), inference(subsumption_resolution, [], [f17997, f3989])).
fof(f3989, plain, (~ (e22 = op2(e22, e24)) | spl42_556), inference(avatar_component_clause, [], [f3988])).
fof(f3988, plain, (spl42_556 <=> (e22 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_556])])).
fof(f17997, plain, ((e22 = op2(e22, e24)) | (~ spl42_470 | ~ spl42_654)), inference(forward_demodulation, [], [f17984, f4396])).
fof(f4396, plain, ((e22 = op2(e20, e24)) | ~ spl42_654), inference(avatar_component_clause, [], [f4394])).
fof(f4394, plain, (spl42_654 <=> (e22 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_654])])).
fof(f17984, plain, ((e22 = op2(op2(e20, e24), e24)) | ~ spl42_470), inference(backward_demodulation, [], [f472, f3634])).
fof(f472, plain, (e22 = op2(op2(op2(e24, e22), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f17942, plain, (~ spl42_516 | spl42_789), inference(avatar_contradiction_clause, [], [f17941])).
fof(f17941, plain, ($false | (~ spl42_516 | spl42_789)), inference(subsumption_resolution, [], [f17938, f5346])).
fof(f5346, plain, (~ (e24 = h4(e15)) | spl42_789), inference(avatar_component_clause, [], [f5344])).
fof(f5344, plain, (spl42_789 <=> (e24 = h4(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_789])])).
fof(f17938, plain, ((e24 = h4(e15)) | ~ spl42_516), inference(backward_demodulation, [], [f1225, f3824])).
fof(f3824, plain, ((e24 = op2(e23, e23)) | ~ spl42_516), inference(avatar_component_clause, [], [f3822])).
fof(f3822, plain, (spl42_516 <=> (e24 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_516])])).
fof(f17915, plain, (~ spl42_388 | ~ spl42_541 | spl42_589), inference(avatar_contradiction_clause, [], [f17914])).
fof(f17914, plain, ($false | (~ spl42_388 | ~ spl42_541 | spl42_589)), inference(subsumption_resolution, [], [f17908, f4126])).
fof(f4126, plain, (~ (e20 = op2(e21, e26)) | spl42_589), inference(avatar_component_clause, [], [f4125])).
fof(f17908, plain, ((e20 = op2(e21, e26)) | (~ spl42_388 | ~ spl42_541)), inference(backward_demodulation, [], [f17190, f3928])).
fof(f17190, plain, ((e20 = op2(op2(e22, e26), e26)) | ~ spl42_388), inference(backward_demodulation, [], [f460, f3294])).
fof(f17898, plain, (~ spl42_552 | ~ spl42_553), inference(avatar_contradiction_clause, [], [f17897])).
fof(f17897, plain, ($false | (~ spl42_552 | ~ spl42_553)), inference(subsumption_resolution, [], [f17896, f1132])).
fof(f17896, plain, ((e25 = e26) | (~ spl42_552 | ~ spl42_553)), inference(forward_demodulation, [], [f3977, f3973])).
fof(f3977, plain, ((e26 = op2(e22, e25)) | ~ spl42_553), inference(avatar_component_clause, [], [f3975])).
fof(f3975, plain, (spl42_553 <=> (e26 = op2(e22, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_553])])).
fof(f17842, plain, (~ spl42_439 | ~ spl42_586), inference(avatar_split_clause, [], [f17832, f4112, f3503])).
fof(f4112, plain, (spl42_586 <=> (e24 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_586])])).
fof(f17832, plain, (~ (e24 = op2(e25, e20)) | ~ spl42_586), inference(backward_demodulation, [], [f810, f4114])).
fof(f4114, plain, ((e24 = op2(e22, e20)) | ~ spl42_586), inference(avatar_component_clause, [], [f4112])).
fof(f810, plain, ~ (op2(e22, e20) = op2(e25, e20)), inference(cnf_transformation, [], [f8])).
fof(f17822, plain, (~ spl42_382 | spl42_529 | ~ spl42_616), inference(avatar_contradiction_clause, [], [f17821])).
fof(f17821, plain, ($false | (~ spl42_382 | spl42_529 | ~ spl42_616)), inference(subsumption_resolution, [], [f17820, f3877])).
fof(f3877, plain, (~ (e23 = op2(e23, e21)) | spl42_529), inference(avatar_component_clause, [], [f3876])).
fof(f3876, plain, (spl42_529 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_529])])).
fof(f17820, plain, ((e23 = op2(e23, e21)) | (~ spl42_382 | ~ spl42_616)), inference(forward_demodulation, [], [f17813, f3269])).
fof(f17813, plain, ((e23 = op2(op2(e26, e21), e21)) | ~ spl42_616), inference(backward_demodulation, [], [f476, f4238])).
fof(f4238, plain, ((e26 = op2(e21, e23)) | ~ spl42_616), inference(avatar_component_clause, [], [f4236])).
fof(f4236, plain, (spl42_616 <=> (e26 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_616])])).
fof(f17802, plain, (~ spl42_630 | spl42_1227), inference(avatar_contradiction_clause, [], [f17801])).
fof(f17801, plain, ($false | (~ spl42_630 | spl42_1227)), inference(subsumption_resolution, [], [f17800, f8020])).
fof(f8020, plain, (~ (e26 = h2(e15)) | spl42_1227), inference(avatar_component_clause, [], [f8018])).
fof(f17800, plain, ((e26 = h2(e15)) | ~ spl42_630), inference(backward_demodulation, [], [f1211, f4296])).
fof(f17736, plain, (~ spl42_670 | ~ spl42_649), inference(avatar_split_clause, [], [f17735, f4373, f4460])).
fof(f4460, plain, (spl42_670 <=> (e24 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_670])])).
fof(f17735, plain, (~ (e24 = op2(e20, e22)) | ~ spl42_649), inference(forward_demodulation, [], [f957, f4375])).
fof(f957, plain, ~ (op2(e20, e22) = op2(e20, e25)), inference(cnf_transformation, [], [f8])).
fof(f17721, plain, (~ spl42_603 | spl42_628 | ~ spl42_674), inference(avatar_contradiction_clause, [], [f17720])).
fof(f17720, plain, ($false | (~ spl42_603 | spl42_628 | ~ spl42_674)), inference(subsumption_resolution, [], [f17719, f4287])).
fof(f4287, plain, (~ (e24 = op2(e21, e21)) | spl42_628), inference(avatar_component_clause, [], [f4286])).
fof(f17719, plain, ((e24 = op2(e21, e21)) | (~ spl42_603 | ~ spl42_674)), inference(forward_demodulation, [], [f17718, f4479])).
fof(f4479, plain, ((e21 = op2(e20, e21)) | ~ spl42_674), inference(avatar_component_clause, [], [f4477])).
fof(f17718, plain, ((e24 = op2(op2(e20, e21), e21)) | ~ spl42_603), inference(forward_demodulation, [], [f483, f4185])).
fof(f4185, plain, ((e20 = op2(e21, e24)) | ~ spl42_603), inference(avatar_component_clause, [], [f4183])).
fof(f4183, plain, (spl42_603 <=> (e20 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_603])])).
fof(f483, plain, (e24 = op2(op2(op2(e21, e24), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f17690, plain, (spl42_417 | ~ spl42_501), inference(avatar_contradiction_clause, [], [f17689])).
fof(f17689, plain, ($false | (spl42_417 | ~ spl42_501)), inference(subsumption_resolution, [], [f17688, f3413])).
fof(f3413, plain, (~ (e23 = op2(e25, e23)) | spl42_417), inference(avatar_component_clause, [], [f3412])).
fof(f17688, plain, ((e23 = op2(e25, e23)) | ~ spl42_501), inference(backward_demodulation, [], [f4707, f17687])).
fof(f17687, plain, ((e25 = op2(h4(e15), e23)) | ~ spl42_501), inference(backward_demodulation, [], [f17686, f1225])).
fof(f17686, plain, ((e25 = op2(op2(e23, e23), e23)) | ~ spl42_501), inference(forward_demodulation, [], [f492, f3762])).
fof(f17633, plain, (~ spl42_395 | ~ spl42_401 | ~ spl42_552), inference(avatar_contradiction_clause, [], [f17632])).
fof(f17632, plain, ($false | (~ spl42_395 | ~ spl42_401 | ~ spl42_552)), inference(subsumption_resolution, [], [f17631, f1122])).
fof(f17631, plain, ((e21 = e26) | (~ spl42_395 | ~ spl42_401 | ~ spl42_552)), inference(forward_demodulation, [], [f17630, f3348])).
fof(f17630, plain, ((e26 = op2(e25, e25)) | (~ spl42_395 | ~ spl42_552)), inference(forward_demodulation, [], [f17629, f3973])).
fof(f17629, plain, ((e26 = op2(op2(e22, e25), e25)) | ~ spl42_395), inference(forward_demodulation, [], [f501, f3323])).
fof(f3323, plain, ((e22 = op2(e25, e26)) | ~ spl42_395), inference(avatar_component_clause, [], [f3321])).
fof(f3321, plain, (spl42_395 <=> (e22 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_395])])).
fof(f17626, plain, (~ spl42_386 | ~ spl42_351), inference(avatar_split_clause, [], [f12692, f3139, f3284])).
fof(f3284, plain, (spl42_386 <=> (e20 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_386])])).
fof(f12692, plain, (~ (e20 = op2(e26, e20)) | ~ spl42_351), inference(forward_demodulation, [], [f1074, f3141])).
fof(f1074, plain, ~ (op2(e26, e20) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f17615, plain, (~ spl42_362 | ~ spl42_369), inference(avatar_split_clause, [], [f13799, f3213, f3184])).
fof(f3184, plain, (spl42_362 <=> (e24 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_362])])).
fof(f13799, plain, (~ (e24 = op2(e26, e24)) | ~ spl42_369), inference(forward_demodulation, [], [f1085, f3215])).
fof(f1085, plain, ~ (op2(e26, e23) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f17614, plain, (~ spl42_360 | ~ spl42_388), inference(avatar_split_clause, [], [f17613, f3292, f3176])).
fof(f3176, plain, (spl42_360 <=> (e22 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_360])])).
fof(f17613, plain, (~ (e22 = op2(e26, e24)) | ~ spl42_388), inference(forward_demodulation, [], [f1073, f3294])).
fof(f17600, plain, (spl42_306 | ~ spl42_8 | ~ spl42_70), inference(avatar_split_clause, [], [f17463, f1876, f1620, f2854])).
fof(f17463, plain, ((e14 = op1(e10, e15)) | (~ spl42_8 | ~ spl42_70)), inference(forward_demodulation, [], [f17460, f1622])).
fof(f17460, plain, ((e14 = op1(op1(e16, e15), e15)) | ~ spl42_70), inference(backward_demodulation, [], [f291, f1878])).
fof(f291, plain, (e14 = op1(op1(op1(e15, e14), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f17586, plain, (~ spl42_249 | ~ spl42_263), inference(avatar_split_clause, [], [f17585, f2676, f2618])).
fof(f2676, plain, (spl42_263 <=> (e13 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_263])])).
fof(f17585, plain, (~ (e13 = op1(e11, e16)) | ~ spl42_263), inference(forward_demodulation, [], [f690, f2678])).
fof(f2678, plain, ((e13 = op1(e11, e14)) | ~ spl42_263), inference(avatar_component_clause, [], [f2676])).
fof(f690, plain, ~ (op1(e11, e14) = op1(e11, e16)), inference(cnf_transformation, [], [f7])).
fof(f17546, plain, (spl42_16 | ~ spl42_70 | ~ spl42_139), inference(avatar_split_clause, [], [f17545, f2162, f1876, f1653])).
fof(f17545, plain, ((e11 = op1(e16, e14)) | (~ spl42_70 | ~ spl42_139)), inference(forward_demodulation, [], [f17544, f1878])).
fof(f17544, plain, ((e11 = op1(op1(e15, e14), e14)) | ~ spl42_139), inference(forward_demodulation, [], [f269, f2164])).
fof(f17537, plain, (~ spl42_286 | ~ spl42_139), inference(avatar_split_clause, [], [f17536, f2162, f2771])).
fof(f2771, plain, (spl42_286 <=> (e15 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_286])])).
fof(f17536, plain, (~ (e15 = op1(e11, e11)) | ~ spl42_139), inference(forward_demodulation, [], [f532, f2164])).
fof(f532, plain, ~ (op1(e11, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f17535, plain, (~ spl42_136 | ~ spl42_87), inference(avatar_split_clause, [], [f15802, f1947, f2150])).
fof(f2150, plain, (spl42_136 <=> (e12 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_136])])).
fof(f15802, plain, (~ (e12 = op1(e14, e11)) | ~ spl42_87), inference(forward_demodulation, [], [f542, f1949])).
fof(f542, plain, ~ (op1(e14, e11) = op1(e15, e11)), inference(cnf_transformation, [], [f7])).
fof(f17531, plain, (~ spl42_131 | ~ spl42_117), inference(avatar_split_clause, [], [f17530, f2071, f2129])).
fof(f2129, plain, (spl42_131 <=> (e14 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_131])])).
fof(f17530, plain, (~ (e14 = op1(e14, e12)) | ~ spl42_117), inference(forward_demodulation, [], [f746, f2073])).
fof(f17499, plain, (~ spl42_79 | ~ spl42_58), inference(avatar_split_clause, [], [f13615, f1827, f1914])).
fof(f1914, plain, (spl42_79 <=> (e11 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_79])])).
fof(f13615, plain, (~ (e11 = op1(e15, e12)) | ~ spl42_58), inference(forward_demodulation, [], [f768, f1829])).
fof(f768, plain, ~ (op1(e15, e12) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f17494, plain, (spl42_52 | ~ spl42_6 | ~ spl42_35), inference(avatar_split_clause, [], [f17473, f1731, f1611, f1802])).
fof(f17473, plain, ((e12 = op1(e15, e16)) | (~ spl42_6 | ~ spl42_35)), inference(forward_demodulation, [], [f17471, f1613])).
fof(f17471, plain, ((e12 = op1(op1(e16, e16), e16)) | ~ spl42_35), inference(backward_demodulation, [], [f278, f1733])).
fof(f1733, plain, ((e16 = op1(e16, e12)) | ~ spl42_35), inference(avatar_component_clause, [], [f1731])).
fof(f278, plain, (e12 = op1(op1(op1(e16, e12), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f17441, plain, (spl42_4 | ~ spl42_26 | ~ spl42_105), inference(avatar_contradiction_clause, [], [f17440])).
fof(f17440, plain, ($false | (spl42_4 | ~ spl42_26 | ~ spl42_105)), inference(subsumption_resolution, [], [f17439, f1604])).
fof(f1604, plain, (~ (e13 = op1(e16, e16)) | spl42_4), inference(avatar_component_clause, [], [f1603])).
fof(f1603, plain, (spl42_4 <=> (e13 = op1(e16, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_4])])).
fof(f17439, plain, ((e13 = op1(e16, e16)) | (~ spl42_26 | ~ spl42_105)), inference(backward_demodulation, [], [f12209, f2023])).
fof(f2023, plain, ((e16 = op1(e14, e16)) | ~ spl42_105), inference(avatar_component_clause, [], [f2021])).
fof(f2021, plain, (spl42_105 <=> (e16 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_105])])).
fof(f17428, plain, (~ spl42_116 | ~ spl42_117), inference(avatar_contradiction_clause, [], [f17427])).
fof(f17427, plain, ($false | (~ spl42_116 | ~ spl42_117)), inference(subsumption_resolution, [], [f17426, f1106])).
fof(f1106, plain, ~ (e13 = e14), inference(cnf_transformation, [], [f9])).
fof(f17426, plain, ((e13 = e14) | (~ spl42_116 | ~ spl42_117)), inference(forward_demodulation, [], [f2073, f2069])).
fof(f2069, plain, ((e13 = op1(e14, e14)) | ~ spl42_116), inference(avatar_component_clause, [], [f2067])).
fof(f2067, plain, (spl42_116 <=> (e13 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_116])])).
fof(f17408, plain, (~ spl42_134 | ~ spl42_139), inference(avatar_contradiction_clause, [], [f17407])).
fof(f17407, plain, ($false | (~ spl42_134 | ~ spl42_139)), inference(subsumption_resolution, [], [f17406, f1095])).
fof(f17406, plain, ((e10 = e15) | (~ spl42_134 | ~ spl42_139)), inference(forward_demodulation, [], [f2164, f2144])).
fof(f2144, plain, ((e10 = op1(e14, e11)) | ~ spl42_134), inference(avatar_component_clause, [], [f2142])).
fof(f2142, plain, (spl42_134 <=> (e10 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_134])])).
fof(f17344, plain, (spl42_33 | ~ spl42_182 | ~ spl42_214), inference(avatar_contradiction_clause, [], [f17343])).
fof(f17343, plain, ($false | (spl42_33 | ~ spl42_182 | ~ spl42_214)), inference(subsumption_resolution, [], [f17342, f1724])).
fof(f1724, plain, (~ (e14 = op1(e16, e12)) | spl42_33), inference(avatar_component_clause, [], [f1723])).
fof(f1723, plain, (spl42_33 <=> (e14 = op1(e16, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_33])])).
fof(f17342, plain, ((e14 = op1(e16, e12)) | (~ spl42_182 | ~ spl42_214)), inference(forward_demodulation, [], [f17335, f2342])).
fof(f17335, plain, ((e14 = op1(op1(e13, e12), e12)) | ~ spl42_214), inference(backward_demodulation, [], [f288, f2475])).
fof(f2475, plain, ((e13 = op1(e12, e14)) | ~ spl42_214), inference(avatar_component_clause, [], [f2473])).
fof(f2473, plain, (spl42_214 <=> (e13 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_214])])).
fof(f288, plain, (e14 = op1(op1(op1(e12, e14), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f17301, plain, (~ spl42_234 | ~ spl42_255), inference(avatar_contradiction_clause, [], [f17300])).
fof(f17300, plain, ($false | (~ spl42_234 | ~ spl42_255)), inference(subsumption_resolution, [], [f17299, f1104])).
fof(f17299, plain, ((e12 = e15) | (~ spl42_234 | ~ spl42_255)), inference(forward_demodulation, [], [f17292, f2558])).
fof(f2558, plain, ((e12 = op1(e12, e11)) | ~ spl42_234), inference(avatar_component_clause, [], [f2556])).
fof(f2556, plain, (spl42_234 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_234])])).
fof(f17292, plain, ((e15 = op1(e12, e11)) | (~ spl42_234 | ~ spl42_255)), inference(backward_demodulation, [], [f13742, f2558])).
fof(f17264, plain, (~ spl42_255 | ~ spl42_258), inference(avatar_contradiction_clause, [], [f17263])).
fof(f17263, plain, ($false | (~ spl42_255 | ~ spl42_258)), inference(subsumption_resolution, [], [f17262, f1104])).
fof(f17262, plain, ((e12 = e15) | (~ spl42_255 | ~ spl42_258)), inference(forward_demodulation, [], [f2657, f2645])).
fof(f2657, plain, ((e15 = op1(e11, e15)) | ~ spl42_258), inference(avatar_component_clause, [], [f2655])).
fof(f2655, plain, (spl42_258 <=> (e15 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_258])])).
fof(f17244, plain, (~ spl42_39 | spl42_186 | ~ spl42_273), inference(avatar_contradiction_clause, [], [f17243])).
fof(f17243, plain, ($false | (~ spl42_39 | spl42_186 | ~ spl42_273)), inference(subsumption_resolution, [], [f17242, f2358])).
fof(f2358, plain, (~ (e13 = op1(e13, e11)) | spl42_186), inference(avatar_component_clause, [], [f2357])).
fof(f2357, plain, (spl42_186 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_186])])).
fof(f17242, plain, ((e13 = op1(e13, e11)) | (~ spl42_39 | ~ spl42_273)), inference(forward_demodulation, [], [f17236, f1750])).
fof(f17236, plain, ((e13 = op1(op1(e16, e11), e11)) | ~ spl42_273), inference(backward_demodulation, [], [f280, f2719])).
fof(f2719, plain, ((e16 = op1(e11, e13)) | ~ spl42_273), inference(avatar_component_clause, [], [f2717])).
fof(f2717, plain, (spl42_273 <=> (e16 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_273])])).
fof(f17222, plain, (~ spl42_8 | spl42_104 | ~ spl42_299), inference(avatar_contradiction_clause, [], [f17221])).
fof(f17221, plain, ($false | (~ spl42_8 | spl42_104 | ~ spl42_299)), inference(subsumption_resolution, [], [f17215, f2018])).
fof(f2018, plain, (~ (e15 = op1(e14, e16)) | spl42_104), inference(avatar_component_clause, [], [f2017])).
fof(f2017, plain, (spl42_104 <=> (e15 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_104])])).
fof(f17215, plain, ((e15 = op1(e14, e16)) | (~ spl42_8 | ~ spl42_299)), inference(backward_demodulation, [], [f14670, f2827])).
fof(f2827, plain, ((e14 = op1(e10, e16)) | ~ spl42_299), inference(avatar_component_clause, [], [f2825])).
fof(f2825, plain, (spl42_299 <=> (e14 = op1(e10, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_299])])).
fof(f17198, plain, (~ spl42_388 | ~ spl42_540 | spl42_638), inference(avatar_contradiction_clause, [], [f17197])).
fof(f17197, plain, ($false | (~ spl42_388 | ~ spl42_540 | spl42_638)), inference(subsumption_resolution, [], [f17196, f4329])).
fof(f4329, plain, (~ (e20 = op2(e20, e26)) | spl42_638), inference(avatar_component_clause, [], [f4328])).
fof(f17196, plain, ((e20 = op2(e20, e26)) | (~ spl42_388 | ~ spl42_540)), inference(forward_demodulation, [], [f17190, f3924])).
fof(f3924, plain, ((e20 = op2(e22, e26)) | ~ spl42_540), inference(avatar_component_clause, [], [f3922])).
fof(f17183, plain, (~ spl42_349 | ~ spl42_399 | ~ spl42_688), inference(avatar_contradiction_clause, [], [f17182])).
fof(f17182, plain, ($false | (~ spl42_349 | ~ spl42_399 | ~ spl42_688)), inference(subsumption_resolution, [], [f17181, f1132])).
fof(f17181, plain, ((e25 = e26) | (~ spl42_349 | ~ spl42_399 | ~ spl42_688)), inference(forward_demodulation, [], [f17176, f3132])).
fof(f17176, plain, ((e26 = op2(e26, e26)) | (~ spl42_399 | ~ spl42_688)), inference(backward_demodulation, [], [f12687, f3339])).
fof(f3339, plain, ((e26 = op2(e25, e26)) | ~ spl42_399), inference(avatar_component_clause, [], [f3337])).
fof(f3337, plain, (spl42_399 <=> (e26 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_399])])).
fof(f17144, plain, (~ spl42_439 | spl42_488 | ~ spl42_657), inference(avatar_contradiction_clause, [], [f17143])).
fof(f17143, plain, ($false | (~ spl42_439 | spl42_488 | ~ spl42_657)), inference(subsumption_resolution, [], [f17138, f3707])).
fof(f3707, plain, (~ (e24 = op2(e24, e20)) | spl42_488), inference(avatar_component_clause, [], [f3706])).
fof(f3706, plain, (spl42_488 <=> (e24 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_488])])).
fof(f17138, plain, ((e24 = op2(e24, e20)) | (~ spl42_439 | ~ spl42_657)), inference(backward_demodulation, [], [f16944, f3505])).
fof(f16944, plain, ((e24 = op2(op2(e25, e20), e20)) | ~ spl42_657), inference(forward_demodulation, [], [f482, f4408])).
fof(f4408, plain, ((e25 = op2(e20, e24)) | ~ spl42_657), inference(avatar_component_clause, [], [f4406])).
fof(f4406, plain, (spl42_657 <=> (e25 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_657])])).
fof(f482, plain, (e24 = op2(op2(op2(e20, e24), e20), e20)), inference(cnf_transformation, [], [f6])).
fof(f17142, plain, (~ spl42_397 | ~ spl42_439), inference(avatar_split_clause, [], [f17135, f3503, f3329])).
fof(f17135, plain, (~ (e24 = op2(e25, e26)) | ~ spl42_439), inference(backward_demodulation, [], [f1054, f3505])).
fof(f1054, plain, ~ (op2(e25, e20) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f17130, plain, (~ spl42_369 | ~ spl42_443 | spl42_592), inference(avatar_contradiction_clause, [], [f17129])).
fof(f17129, plain, ($false | (~ spl42_369 | ~ spl42_443 | spl42_592)), inference(subsumption_resolution, [], [f17124, f4138])).
fof(f4138, plain, (~ (e23 = op2(e21, e26)) | spl42_592), inference(avatar_component_clause, [], [f4137])).
fof(f17124, plain, ((e23 = op2(e21, e26)) | (~ spl42_369 | ~ spl42_443)), inference(backward_demodulation, [], [f12623, f3522])).
fof(f17128, plain, (~ spl42_443 | ~ spl42_603 | ~ spl42_657), inference(avatar_contradiction_clause, [], [f17127])).
fof(f17127, plain, ($false | (~ spl42_443 | ~ spl42_603 | ~ spl42_657)), inference(subsumption_resolution, [], [f17126, f1132])).
fof(f17126, plain, ((e25 = e26) | (~ spl42_443 | ~ spl42_603 | ~ spl42_657)), inference(forward_demodulation, [], [f17125, f4408])).
fof(f17125, plain, ((e26 = op2(e20, e24)) | (~ spl42_443 | ~ spl42_603)), inference(forward_demodulation, [], [f17116, f4185])).
fof(f17116, plain, ((e26 = op2(op2(e21, e24), e24)) | ~ spl42_443), inference(backward_demodulation, [], [f500, f3522])).
fof(f17115, plain, (~ spl42_462 | spl42_1062), inference(avatar_contradiction_clause, [], [f17114])).
fof(f17114, plain, ($false | (~ spl42_462 | spl42_1062)), inference(subsumption_resolution, [], [f17113, f6976])).
fof(f6976, plain, (~ (e26 = h5(e15)) | spl42_1062), inference(avatar_component_clause, [], [f6974])).
fof(f17113, plain, ((e26 = h5(e15)) | ~ spl42_462), inference(backward_demodulation, [], [f1232, f3600])).
fof(f3600, plain, ((e26 = op2(e24, e24)) | ~ spl42_462), inference(avatar_component_clause, [], [f3598])).
fof(f17061, plain, (~ spl42_518 | spl42_1117), inference(avatar_contradiction_clause, [], [f17060])).
fof(f17060, plain, ($false | (~ spl42_518 | spl42_1117)), inference(subsumption_resolution, [], [f17059, f7324])).
fof(f7324, plain, (~ (e26 = h4(e15)) | spl42_1117), inference(avatar_component_clause, [], [f7322])).
fof(f17059, plain, ((e26 = h4(e15)) | ~ spl42_518), inference(backward_demodulation, [], [f1225, f3832])).
fof(f3832, plain, ((e26 = op2(e23, e23)) | ~ spl42_518), inference(avatar_component_clause, [], [f3830])).
fof(f3830, plain, (spl42_518 <=> (e26 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_518])])).
fof(f17017, plain, (~ spl42_574 | spl42_1172), inference(avatar_contradiction_clause, [], [f17016])).
fof(f17016, plain, ($false | (~ spl42_574 | spl42_1172)), inference(subsumption_resolution, [], [f17015, f7672])).
fof(f7672, plain, (~ (e26 = h3(e15)) | spl42_1172), inference(avatar_component_clause, [], [f7670])).
fof(f17015, plain, ((e26 = h3(e15)) | ~ spl42_574), inference(backward_demodulation, [], [f1218, f4064])).
fof(f4064, plain, ((e26 = op2(e22, e22)) | ~ spl42_574), inference(avatar_component_clause, [], [f4062])).
fof(f4062, plain, (spl42_574 <=> (e26 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_574])])).
fof(f16989, plain, (~ spl42_597 | ~ spl42_598), inference(avatar_contradiction_clause, [], [f16988])).
fof(f16988, plain, ($false | (~ spl42_597 | ~ spl42_598)), inference(subsumption_resolution, [], [f16987, f1118])).
fof(f16987, plain, ((e21 = e22) | (~ spl42_597 | ~ spl42_598)), inference(backward_demodulation, [], [f4164, f4160])).
fof(f4160, plain, ((e21 = op2(e21, e25)) | ~ spl42_597), inference(avatar_component_clause, [], [f4158])).
fof(f4158, plain, (spl42_597 <=> (e21 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_597])])).
fof(f16968, plain, (~ spl42_629 | spl42_852), inference(avatar_contradiction_clause, [], [f16967])).
fof(f16967, plain, ($false | (~ spl42_629 | spl42_852)), inference(subsumption_resolution, [], [f16966, f5663])).
fof(f5663, plain, (~ (e25 = h2(e15)) | spl42_852), inference(avatar_component_clause, [], [f5661])).
fof(f16966, plain, ((e25 = h2(e15)) | ~ spl42_629), inference(backward_demodulation, [], [f1211, f4292])).
fof(f4292, plain, ((e25 = op2(e21, e21)) | ~ spl42_629), inference(avatar_component_clause, [], [f4290])).
fof(f4290, plain, (spl42_629 <=> (e25 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_629])])).
fof(f16929, plain, (spl42_729 | ~ spl42_618 | ~ spl42_733), inference(avatar_split_clause, [], [f16928, f5063, f4245, f5043])).
fof(f4245, plain, (spl42_618 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_618])])).
fof(f16928, plain, ((e22 = h6(e12)) | (~ spl42_618 | ~ spl42_733)), inference(backward_demodulation, [], [f14048, f16927])).
fof(f16927, plain, ((e22 = op2(h2(e15), e21)) | ~ spl42_618), inference(backward_demodulation, [], [f16093, f1211])).
fof(f16093, plain, ((e22 = op2(op2(e21, e21), e21)) | ~ spl42_618), inference(backward_demodulation, [], [f469, f4247])).
fof(f4247, plain, ((e21 = op2(e21, e22)) | ~ spl42_618), inference(avatar_component_clause, [], [f4245])).
fof(f16916, plain, (spl42_597 | ~ spl42_401 | ~ spl42_433 | ~ spl42_733), inference(avatar_split_clause, [], [f16282, f5063, f3478, f3346, f4158])).
fof(f3478, plain, (spl42_433 <=> (e25 = op2(e25, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_433])])).
fof(f16282, plain, ((e21 = op2(e21, e25)) | (~ spl42_401 | ~ spl42_433 | ~ spl42_733)), inference(forward_demodulation, [], [f16278, f3348])).
fof(f16278, plain, ((e21 = op2(op2(e25, e25), e25)) | (~ spl42_433 | ~ spl42_733)), inference(backward_demodulation, [], [f13887, f16275])).
fof(f16275, plain, ((e25 = h6(e10)) | (~ spl42_433 | ~ spl42_733)), inference(backward_demodulation, [], [f13004, f3480])).
fof(f3480, plain, ((e25 = op2(e25, e21)) | ~ spl42_433), inference(avatar_component_clause, [], [f3478])).
fof(f13887, plain, ((e21 = op2(op2(h6(e10), e25), e25)) | ~ spl42_733), inference(forward_demodulation, [], [f466, f13004])).
fof(f466, plain, (e21 = op2(op2(op2(e25, e21), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f16907, plain, (~ spl42_587 | ~ spl42_552), inference(avatar_split_clause, [], [f16906, f3971, f4116])).
fof(f4116, plain, (spl42_587 <=> (e25 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_587])])).
fof(f16906, plain, (~ (e25 = op2(e22, e20)) | ~ spl42_552), inference(forward_demodulation, [], [f990, f3973])).
fof(f990, plain, ~ (op2(e22, e20) = op2(e22, e25)), inference(cnf_transformation, [], [f8])).
fof(f16871, plain, (~ spl42_503 | ~ spl42_552), inference(avatar_split_clause, [], [f16005, f3971, f3768])).
fof(f3768, plain, (spl42_503 <=> (e25 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_503])])).
fof(f16005, plain, (~ (e25 = op2(e23, e25)) | ~ spl42_552), inference(backward_demodulation, [], [f913, f3973])).
fof(f913, plain, ~ (op2(e22, e25) = op2(e23, e25)), inference(cnf_transformation, [], [f8])).
fof(f16795, plain, (~ spl42_415 | ~ spl42_733), inference(avatar_split_clause, [], [f13876, f5063, f3404])).
fof(f3404, plain, (spl42_415 <=> (e21 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_415])])).
fof(f13876, plain, (~ (e21 = op2(e25, e23)) | ~ spl42_733), inference(forward_demodulation, [], [f4759, f5064])).
fof(f4759, plain, ~ (op2(e25, e23) = h6(e15)), inference(backward_demodulation, [], [f1065, f1239])).
fof(f1065, plain, ~ (op2(e25, e23) = op2(e25, e25)), inference(cnf_transformation, [], [f8])).
fof(f16794, plain, (spl42_395 | ~ spl42_349 | ~ spl42_378), inference(avatar_split_clause, [], [f16354, f3250, f3130, f3321])).
fof(f3250, plain, (spl42_378 <=> (e26 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_378])])).
fof(f16354, plain, ((e22 = op2(e25, e26)) | (~ spl42_349 | ~ spl42_378)), inference(forward_demodulation, [], [f16350, f3132])).
fof(f16350, plain, ((e22 = op2(op2(e26, e26), e26)) | ~ spl42_378), inference(backward_demodulation, [], [f474, f3252])).
fof(f3252, plain, ((e26 = op2(e26, e22)) | ~ spl42_378), inference(avatar_component_clause, [], [f3250])).
fof(f474, plain, (e22 = op2(op2(op2(e26, e22), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f16743, plain, (spl42_215 | ~ spl42_131), inference(avatar_split_clause, [], [f16571, f2129, f2477])).
fof(f2477, plain, (spl42_215 <=> (e14 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_215])])).
fof(f16571, plain, ((e14 = op1(e12, e14)) | ~ spl42_131), inference(backward_demodulation, [], [f290, f16564])).
fof(f16564, plain, ((e12 = op1(op1(e14, e14), e14)) | ~ spl42_131), inference(backward_demodulation, [], [f276, f2131])).
fof(f2131, plain, ((e14 = op1(e14, e12)) | ~ spl42_131), inference(avatar_component_clause, [], [f2129])).
fof(f290, plain, (e14 = op1(op1(op1(e14, e14), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f16721, plain, (spl42_296 | ~ spl42_39 | ~ spl42_148), inference(avatar_split_clause, [], [f16720, f2200, f1748, f2813])).
fof(f2813, plain, (spl42_296 <=> (e11 = op1(e10, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_296])])).
fof(f2200, plain, (spl42_148 <=> (e10 = op1(e13, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_148])])).
fof(f16720, plain, ((e11 = op1(e10, e16)) | (~ spl42_39 | ~ spl42_148)), inference(forward_demodulation, [], [f12233, f2202])).
fof(f2202, plain, ((e10 = op1(e13, e16)) | ~ spl42_148), inference(avatar_component_clause, [], [f2200])).
fof(f16699, plain, (~ spl42_111 | ~ spl42_209), inference(avatar_split_clause, [], [f14446, f2452, f2046])).
fof(f2046, plain, (spl42_111 <=> (e15 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_111])])).
fof(f14446, plain, (~ (e15 = op1(e14, e15)) | ~ spl42_209), inference(backward_demodulation, [], [f620, f2454])).
fof(f620, plain, ~ (op1(e12, e15) = op1(e14, e15)), inference(cnf_transformation, [], [f7])).
fof(f16647, plain, (~ spl42_45 | spl42_197 | ~ spl42_199), inference(avatar_contradiction_clause, [], [f16646])).
fof(f16646, plain, ($false | (~ spl42_45 | spl42_197 | ~ spl42_199)), inference(subsumption_resolution, [], [f16645, f2404])).
fof(f2404, plain, (~ (e10 = op1(e12, e16)) | spl42_197), inference(avatar_component_clause, [], [f2403])).
fof(f2403, plain, (spl42_197 <=> (e10 = op1(e12, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_197])])).
fof(f16645, plain, ((e10 = op1(e12, e16)) | (~ spl42_45 | ~ spl42_199)), inference(forward_demodulation, [], [f16638, f2413])).
fof(f16638, plain, ((e10 = op1(op1(e12, e16), e16)) | ~ spl42_45), inference(backward_demodulation, [], [f264, f1775])).
fof(f1775, plain, ((e12 = op1(e16, e10)) | ~ spl42_45), inference(avatar_component_clause, [], [f1773])).
fof(f1773, plain, (spl42_45 <=> (e12 = op1(e16, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_45])])).
fof(f264, plain, (e10 = op1(op1(op1(e16, e10), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f16632, plain, (~ spl42_58 | ~ spl42_61), inference(avatar_contradiction_clause, [], [f16631])).
fof(f16631, plain, ($false | (~ spl42_58 | ~ spl42_61)), inference(subsumption_resolution, [], [f16630, f1099])).
fof(f1099, plain, ~ (e11 = e14), inference(cnf_transformation, [], [f9])).
fof(f16630, plain, ((e11 = e14) | (~ spl42_58 | ~ spl42_61)), inference(forward_demodulation, [], [f1841, f1829])).
fof(f1841, plain, ((e14 = op1(e15, e15)) | ~ spl42_61), inference(avatar_component_clause, [], [f1839])).
fof(f1839, plain, (spl42_61 <=> (e14 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_61])])).
fof(f16523, plain, (spl42_29 | ~ spl42_182 | ~ spl42_242), inference(avatar_contradiction_clause, [], [f16522])).
fof(f16522, plain, ($false | (spl42_29 | ~ spl42_182 | ~ spl42_242)), inference(subsumption_resolution, [], [f16518, f1708])).
fof(f1708, plain, (~ (e10 = op1(e16, e12)) | spl42_29), inference(avatar_component_clause, [], [f1707])).
fof(f1707, plain, (spl42_29 <=> (e10 = op1(e16, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_29])])).
fof(f16518, plain, ((e10 = op1(e16, e12)) | (~ spl42_182 | ~ spl42_242)), inference(backward_demodulation, [], [f16431, f2342])).
fof(f16431, plain, ((e10 = op1(op1(e13, e12), e12)) | ~ spl42_242), inference(backward_demodulation, [], [f260, f2591])).
fof(f2591, plain, ((e13 = op1(e12, e10)) | ~ spl42_242), inference(avatar_component_clause, [], [f2589])).
fof(f2589, plain, (spl42_242 <=> (e13 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_242])])).
fof(f16520, plain, (~ spl42_133 | ~ spl42_182), inference(avatar_split_clause, [], [f16514, f2340, f2137])).
fof(f16514, plain, (~ (e16 = op1(e14, e12)) | ~ spl42_182), inference(backward_demodulation, [], [f560, f2342])).
fof(f560, plain, ~ (op1(e13, e12) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f16493, plain, (spl42_133 | ~ spl42_199 | ~ spl42_229), inference(avatar_split_clause, [], [f16492, f2535, f2411, f2137])).
fof(f2535, plain, (spl42_229 <=> (e14 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_229])])).
fof(f16492, plain, ((e16 = op1(e14, e12)) | (~ spl42_199 | ~ spl42_229)), inference(forward_demodulation, [], [f16485, f2537])).
fof(f2537, plain, ((e14 = op1(e12, e12)) | ~ spl42_229), inference(avatar_component_clause, [], [f2535])).
fof(f16485, plain, ((e16 = op1(op1(e12, e12), e12)) | ~ spl42_199), inference(backward_demodulation, [], [f302, f2413])).
fof(f16458, plain, (~ spl42_39 | ~ spl42_238 | ~ spl42_255), inference(avatar_contradiction_clause, [], [f16457])).
fof(f16457, plain, ($false | (~ spl42_39 | ~ spl42_238 | ~ spl42_255)), inference(subsumption_resolution, [], [f16456, f1107])).
fof(f1107, plain, ~ (e13 = e15), inference(cnf_transformation, [], [f9])).
fof(f16456, plain, ((e13 = e15) | (~ spl42_39 | ~ spl42_238 | ~ spl42_255)), inference(forward_demodulation, [], [f16450, f1750])).
fof(f16450, plain, ((e15 = op1(e16, e11)) | (~ spl42_238 | ~ spl42_255)), inference(backward_demodulation, [], [f13742, f2574])).
fof(f2574, plain, ((e16 = op1(e12, e11)) | ~ spl42_238), inference(avatar_component_clause, [], [f2572])).
fof(f2572, plain, (spl42_238 <=> (e16 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_238])])).
fof(f16430, plain, (~ spl42_30 | spl42_150 | ~ spl42_249), inference(avatar_contradiction_clause, [], [f16429])).
fof(f16429, plain, ($false | (~ spl42_30 | spl42_150 | ~ spl42_249)), inference(subsumption_resolution, [], [f16424, f2209])).
fof(f2209, plain, (~ (e12 = op1(e13, e16)) | spl42_150), inference(avatar_component_clause, [], [f2208])).
fof(f16424, plain, ((e12 = op1(e13, e16)) | (~ spl42_30 | ~ spl42_249)), inference(backward_demodulation, [], [f15663, f2620])).
fof(f15663, plain, ((e12 = op1(op1(e11, e16), e16)) | ~ spl42_30), inference(backward_demodulation, [], [f278, f1713])).
fof(f1713, plain, ((e11 = op1(e16, e12)) | ~ spl42_30), inference(avatar_component_clause, [], [f1711])).
fof(f1711, plain, (spl42_30 <=> (e11 = op1(e16, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_30])])).
fof(f16408, plain, (~ spl42_87 | spl42_235 | ~ spl42_272), inference(avatar_contradiction_clause, [], [f16407])).
fof(f16407, plain, ($false | (~ spl42_87 | spl42_235 | ~ spl42_272)), inference(subsumption_resolution, [], [f16406, f2561])).
fof(f2561, plain, (~ (e13 = op1(e12, e11)) | spl42_235), inference(avatar_component_clause, [], [f2560])).
fof(f2560, plain, (spl42_235 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_235])])).
fof(f16406, plain, ((e13 = op1(e12, e11)) | (~ spl42_87 | ~ spl42_272)), inference(forward_demodulation, [], [f16401, f1949])).
fof(f16401, plain, ((e13 = op1(op1(e15, e11), e11)) | ~ spl42_272), inference(backward_demodulation, [], [f280, f2715])).
fof(f2715, plain, ((e15 = op1(e11, e13)) | ~ spl42_272), inference(avatar_component_clause, [], [f2713])).
fof(f2713, plain, (spl42_272 <=> (e15 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_272])])).
fof(f16399, plain, (~ spl42_8 | spl42_251 | ~ spl42_296), inference(avatar_contradiction_clause, [], [f16398])).
fof(f16398, plain, ($false | (~ spl42_8 | spl42_251 | ~ spl42_296)), inference(subsumption_resolution, [], [f16392, f2627])).
fof(f2627, plain, (~ (e15 = op1(e11, e16)) | spl42_251), inference(avatar_component_clause, [], [f2626])).
fof(f2626, plain, (spl42_251 <=> (e15 = op1(e11, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_251])])).
fof(f16392, plain, ((e15 = op1(e11, e16)) | (~ spl42_8 | ~ spl42_296)), inference(backward_demodulation, [], [f14670, f2815])).
fof(f2815, plain, ((e11 = op1(e10, e16)) | ~ spl42_296), inference(avatar_component_clause, [], [f2813])).
fof(f16379, plain, (~ spl42_346 | ~ spl42_349), inference(avatar_contradiction_clause, [], [f16378])).
fof(f16378, plain, ($false | (~ spl42_346 | ~ spl42_349)), inference(subsumption_resolution, [], [f16376, f1125])).
fof(f16376, plain, ((e22 = e25) | (~ spl42_346 | ~ spl42_349)), inference(backward_demodulation, [], [f3132, f3120])).
fof(f3120, plain, ((e22 = op2(e26, e26)) | ~ spl42_346), inference(avatar_component_clause, [], [f3118])).
fof(f3118, plain, (spl42_346 <=> (e22 = op2(e26, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_346])])).
fof(f16374, plain, (~ spl42_351 | ~ spl42_355), inference(avatar_contradiction_clause, [], [f16373])).
fof(f16373, plain, ($false | (~ spl42_351 | ~ spl42_355)), inference(subsumption_resolution, [], [f16372, f1115])).
fof(f16372, plain, ((e20 = e24) | (~ spl42_351 | ~ spl42_355)), inference(forward_demodulation, [], [f3157, f3141])).
fof(f3157, plain, ((e24 = op2(e26, e25)) | ~ spl42_355), inference(avatar_component_clause, [], [f3155])).
fof(f3155, plain, (spl42_355 <=> (e24 = op2(e26, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_355])])).
fof(f16361, plain, (~ spl42_366 | ~ spl42_369), inference(avatar_contradiction_clause, [], [f16360])).
fof(f16360, plain, ($false | (~ spl42_366 | ~ spl42_369)), inference(subsumption_resolution, [], [f16359, f1120])).
fof(f16359, plain, ((e21 = e24) | (~ spl42_366 | ~ spl42_369)), inference(backward_demodulation, [], [f3215, f3203])).
fof(f3203, plain, ((e21 = op2(e26, e23)) | ~ spl42_366), inference(avatar_component_clause, [], [f3201])).
fof(f3201, plain, (spl42_366 <=> (e21 = op2(e26, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_366])])).
fof(f16312, plain, (~ spl42_351 | ~ spl42_420 | spl42_648), inference(avatar_contradiction_clause, [], [f16311])).
fof(f16311, plain, ($false | (~ spl42_351 | ~ spl42_420 | spl42_648)), inference(subsumption_resolution, [], [f16310, f4370])).
fof(f4370, plain, (~ (e23 = op2(e20, e25)) | spl42_648), inference(avatar_component_clause, [], [f4369])).
fof(f16310, plain, ((e23 = op2(e20, e25)) | (~ spl42_351 | ~ spl42_420)), inference(forward_demodulation, [], [f16306, f3141])).
fof(f16306, plain, ((e23 = op2(op2(e26, e25), e25)) | ~ spl42_420), inference(backward_demodulation, [], [f480, f3426])).
fof(f16220, plain, (spl42_782 | ~ spl42_517), inference(avatar_split_clause, [], [f16219, f3826, f5309])).
fof(f3826, plain, (spl42_517 <=> (e25 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_517])])).
fof(f16219, plain, ((e25 = h4(e15)) | ~ spl42_517), inference(backward_demodulation, [], [f1225, f3828])).
fof(f3828, plain, ((e25 = op2(e23, e23)) | ~ spl42_517), inference(avatar_component_clause, [], [f3826])).
fof(f16052, plain, (spl42_627 | ~ spl42_610 | ~ spl42_674), inference(avatar_split_clause, [], [f15269, f4477, f4212, f4282])).
fof(f4282, plain, (spl42_627 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_627])])).
fof(f4212, plain, (spl42_610 <=> (e20 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_610])])).
fof(f15269, plain, ((e23 = op2(e21, e21)) | (~ spl42_610 | ~ spl42_674)), inference(forward_demodulation, [], [f15261, f4479])).
fof(f15261, plain, ((e23 = op2(op2(e20, e21), e21)) | ~ spl42_610), inference(backward_demodulation, [], [f476, f4214])).
fof(f4214, plain, ((e20 = op2(e21, e23)) | ~ spl42_610), inference(avatar_component_clause, [], [f4212])).
fof(f16024, plain, (~ spl42_566 | ~ spl42_552), inference(avatar_split_clause, [], [f16023, f3971, f4029])).
fof(f16023, plain, (~ (e25 = op2(e22, e23)) | ~ spl42_552), inference(forward_demodulation, [], [f1002, f3973])).
fof(f1002, plain, ~ (op2(e22, e23) = op2(e22, e25)), inference(cnf_transformation, [], [f8])).
fof(f16008, plain, (~ spl42_580 | ~ spl42_552), inference(avatar_split_clause, [], [f16007, f3971, f4087])).
fof(f4087, plain, (spl42_580 <=> (e25 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_580])])).
fof(f16007, plain, (~ (e25 = op2(e22, e21)) | ~ spl42_552), inference(forward_demodulation, [], [f995, f3973])).
fof(f995, plain, ~ (op2(e22, e21) = op2(e22, e25)), inference(cnf_transformation, [], [f8])).
fof(f15976, plain, (~ spl42_500 | ~ spl42_598), inference(avatar_split_clause, [], [f15975, f4162, f3756])).
fof(f3756, plain, (spl42_500 <=> (e22 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_500])])).
fof(f15975, plain, (~ (e22 = op2(e23, e25)) | ~ spl42_598), inference(forward_demodulation, [], [f909, f4164])).
fof(f909, plain, ~ (op2(e21, e25) = op2(e23, e25)), inference(cnf_transformation, [], [f8])).
fof(f15970, plain, (spl42_494 | ~ spl42_369 | ~ spl42_445), inference(avatar_split_clause, [], [f15387, f3528, f3213, f3731])).
fof(f15387, plain, ((e23 = op2(e23, e26)) | (~ spl42_369 | ~ spl42_445)), inference(backward_demodulation, [], [f12623, f3530])).
fof(f3530, plain, ((e23 = op2(e24, e26)) | ~ spl42_445), inference(avatar_component_clause, [], [f3528])).
fof(f15909, plain, (~ spl42_351 | ~ spl42_441 | spl42_645), inference(avatar_contradiction_clause, [], [f15908])).
fof(f15908, plain, ($false | (~ spl42_351 | ~ spl42_441 | spl42_645)), inference(subsumption_resolution, [], [f15907, f4358])).
fof(f4358, plain, (~ (e20 = op2(e20, e25)) | spl42_645), inference(avatar_component_clause, [], [f4357])).
fof(f4357, plain, (spl42_645 <=> (e20 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_645])])).
fof(f15907, plain, ((e20 = op2(e20, e25)) | (~ spl42_351 | ~ spl42_441)), inference(forward_demodulation, [], [f15906, f3141])).
fof(f15906, plain, ((e20 = op2(op2(e26, e25), e25)) | ~ spl42_441), inference(forward_demodulation, [], [f459, f3513])).
fof(f3513, plain, ((e26 = op2(e25, e20)) | ~ spl42_441), inference(avatar_component_clause, [], [f3511])).
fof(f3511, plain, (spl42_441 <=> (e26 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_441])])).
fof(f15888, plain, (~ spl42_389 | ~ spl42_698), inference(avatar_split_clause, [], [f15887, f4872, f3296])).
fof(f3296, plain, (spl42_389 <=> (e23 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_389])])).
fof(f15887, plain, (~ (e23 = op2(e26, e20)) | ~ spl42_698), inference(forward_demodulation, [], [f4796, f4873])).
fof(f4796, plain, ~ (op2(e26, e20) = h7(e13)), inference(backward_demodulation, [], [f1070, f4789])).
fof(f1070, plain, ~ (op2(e26, e20) = op2(e26, e21)), inference(cnf_transformation, [], [f8])).
fof(f15883, plain, (~ spl42_300 | ~ spl42_6), inference(avatar_split_clause, [], [f13760, f1611, f2829])).
fof(f2829, plain, (spl42_300 <=> (e15 = op1(e10, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_300])])).
fof(f13760, plain, (~ (e15 = op1(e10, e16)) | ~ spl42_6), inference(forward_demodulation, [], [f634, f1613])).
fof(f634, plain, ~ (op1(e10, e16) = op1(e16, e16)), inference(cnf_transformation, [], [f7])).
fof(f15871, plain, (~ spl42_39 | spl42_187 | ~ spl42_266), inference(avatar_contradiction_clause, [], [f15870])).
fof(f15870, plain, ($false | (~ spl42_39 | spl42_187 | ~ spl42_266)), inference(subsumption_resolution, [], [f15869, f2362])).
fof(f2362, plain, (~ (e14 = op1(e13, e11)) | spl42_187), inference(avatar_component_clause, [], [f2361])).
fof(f15869, plain, ((e14 = op1(e13, e11)) | (~ spl42_39 | ~ spl42_266)), inference(forward_demodulation, [], [f15862, f1750])).
fof(f15862, plain, ((e14 = op1(op1(e16, e11), e11)) | ~ spl42_266), inference(backward_demodulation, [], [f287, f2690])).
fof(f2690, plain, ((e16 = op1(e11, e14)) | ~ spl42_266), inference(avatar_component_clause, [], [f2688])).
fof(f2688, plain, (spl42_266 <=> (e16 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_266])])).
fof(f15853, plain, (~ spl42_244 | ~ spl42_209), inference(avatar_split_clause, [], [f15852, f2452, f2597])).
fof(f2597, plain, (spl42_244 <=> (e15 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_244])])).
fof(f15852, plain, (~ (e15 = op1(e12, e10)) | ~ spl42_209), inference(forward_demodulation, [], [f696, f2454])).
fof(f696, plain, ~ (op1(e12, e10) = op1(e12, e15)), inference(cnf_transformation, [], [f7])).
fof(f15820, plain, (~ spl42_180 | ~ spl42_152), inference(avatar_split_clause, [], [f15819, f2216, f2332])).
fof(f2332, plain, (spl42_180 <=> (e14 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_180])])).
fof(f15819, plain, (~ (e14 = op1(e13, e12)) | ~ spl42_152), inference(forward_demodulation, [], [f727, f2218])).
fof(f727, plain, ~ (op1(e13, e12) = op1(e13, e16)), inference(cnf_transformation, [], [f7])).
fof(f15776, plain, (spl42_224 | ~ spl42_122 | ~ spl42_152), inference(avatar_split_clause, [], [f15775, f2216, f2092, f2514])).
fof(f15775, plain, ((e16 = op1(e12, e13)) | (~ spl42_122 | ~ spl42_152)), inference(forward_demodulation, [], [f15580, f2094])).
fof(f15580, plain, ((e16 = op1(op1(e14, e13), e13)) | ~ spl42_152), inference(backward_demodulation, [], [f303, f2218])).
fof(f303, plain, (e16 = op1(op1(op1(e13, e16), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f15763, plain, (~ spl42_108 | ~ spl42_255), inference(avatar_split_clause, [], [f13739, f2643, f2034])).
fof(f2034, plain, (spl42_108 <=> (e12 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_108])])).
fof(f13739, plain, (~ (e12 = op1(e14, e15)) | ~ spl42_255), inference(forward_demodulation, [], [f616, f2645])).
fof(f616, plain, ~ (op1(e11, e15) = op1(e14, e15)), inference(cnf_transformation, [], [f7])).
fof(f15755, plain, (~ spl42_58 | ~ spl42_69 | ~ spl42_255), inference(avatar_contradiction_clause, [], [f15754])).
fof(f15754, plain, ($false | (~ spl42_58 | ~ spl42_69 | ~ spl42_255)), inference(subsumption_resolution, [], [f15753, f1103])).
fof(f15753, plain, ((e12 = e14) | (~ spl42_58 | ~ spl42_69 | ~ spl42_255)), inference(forward_demodulation, [], [f15752, f2645])).
fof(f15752, plain, ((e14 = op1(e11, e15)) | (~ spl42_58 | ~ spl42_69)), inference(forward_demodulation, [], [f15751, f1829])).
fof(f15751, plain, ((e14 = op1(op1(e15, e15), e15)) | ~ spl42_69), inference(forward_demodulation, [], [f291, f1874])).
fof(f1874, plain, ((e15 = op1(e15, e14)) | ~ spl42_69), inference(avatar_component_clause, [], [f1872])).
fof(f1872, plain, (spl42_69 <=> (e15 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_69])])).
fof(f15705, plain, (~ spl42_3 | ~ spl42_6), inference(avatar_contradiction_clause, [], [f15704])).
fof(f15704, plain, ($false | (~ spl42_3 | ~ spl42_6)), inference(subsumption_resolution, [], [f15703, f1104])).
fof(f15703, plain, ((e12 = e15) | (~ spl42_3 | ~ spl42_6)), inference(forward_demodulation, [], [f1613, f1601])).
fof(f1601, plain, ((e12 = op1(e16, e16)) | ~ spl42_3), inference(avatar_component_clause, [], [f1599])).
fof(f1599, plain, (spl42_3 <=> (e12 = op1(e16, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_3])])).
fof(f15699, plain, (~ spl42_4 | ~ spl42_6), inference(avatar_contradiction_clause, [], [f15698])).
fof(f15698, plain, ($false | (~ spl42_4 | ~ spl42_6)), inference(subsumption_resolution, [], [f15697, f1107])).
fof(f15697, plain, ((e13 = e15) | (~ spl42_4 | ~ spl42_6)), inference(backward_demodulation, [], [f1613, f1605])).
fof(f1605, plain, ((e13 = op1(e16, e16)) | ~ spl42_4), inference(avatar_component_clause, [], [f1603])).
fof(f15691, plain, (~ spl42_8 | ~ spl42_14), inference(avatar_contradiction_clause, [], [f15690])).
fof(f15690, plain, ($false | (~ spl42_8 | ~ spl42_14)), inference(subsumption_resolution, [], [f15689, f1096])).
fof(f1096, plain, ~ (e10 = e16), inference(cnf_transformation, [], [f9])).
fof(f15689, plain, ((e10 = e16) | (~ spl42_8 | ~ spl42_14)), inference(forward_demodulation, [], [f1646, f1622])).
fof(f1646, plain, ((e16 = op1(e16, e15)) | ~ spl42_14), inference(avatar_component_clause, [], [f1644])).
fof(f1644, plain, (spl42_14 <=> (e16 = op1(e16, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_14])])).
fof(f15533, plain, (~ spl42_207 | ~ spl42_209), inference(avatar_contradiction_clause, [], [f15532])).
fof(f15532, plain, ($false | (~ spl42_207 | ~ spl42_209)), inference(subsumption_resolution, [], [f15531, f1107])).
fof(f15531, plain, ((e13 = e15) | (~ spl42_207 | ~ spl42_209)), inference(backward_demodulation, [], [f2454, f2446])).
fof(f2446, plain, ((e13 = op1(e12, e15)) | ~ spl42_207), inference(avatar_component_clause, [], [f2444])).
fof(f2444, plain, (spl42_207 <=> (e13 = op1(e12, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_207])])).
fof(f15522, plain, (spl42_90 | ~ spl42_237 | ~ spl42_255), inference(avatar_split_clause, [], [f15515, f2643, f2568, f1959])).
fof(f1959, plain, (spl42_90 <=> (e15 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_90])])).
fof(f2568, plain, (spl42_237 <=> (e15 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_237])])).
fof(f15515, plain, ((e15 = op1(e15, e11)) | (~ spl42_237 | ~ spl42_255)), inference(backward_demodulation, [], [f13742, f2570])).
fof(f2570, plain, ((e15 = op1(e12, e11)) | ~ spl42_237), inference(avatar_component_clause, [], [f2568])).
fof(f15506, plain, (~ spl42_248 | ~ spl42_255), inference(avatar_contradiction_clause, [], [f15505])).
fof(f15505, plain, ($false | (~ spl42_248 | ~ spl42_255)), inference(subsumption_resolution, [], [f15504, f1111])).
fof(f1111, plain, ~ (e15 = e16), inference(cnf_transformation, [], [f9])).
fof(f15504, plain, ((e15 = e16) | (~ spl42_248 | ~ spl42_255)), inference(forward_demodulation, [], [f15496, f13742])).
fof(f15496, plain, ((e16 = op1(op1(e12, e11), e11)) | ~ spl42_248), inference(backward_demodulation, [], [f301, f2616])).
fof(f2616, plain, ((e12 = op1(e11, e16)) | ~ spl42_248), inference(avatar_component_clause, [], [f2614])).
fof(f2614, plain, (spl42_248 <=> (e12 = op1(e11, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_248])])).
fof(f15494, plain, (~ spl42_267 | spl42_284 | ~ spl42_331), inference(avatar_contradiction_clause, [], [f15493])).
fof(f15493, plain, ($false | (~ spl42_267 | spl42_284 | ~ spl42_331)), inference(subsumption_resolution, [], [f15492, f2764])).
fof(f2764, plain, (~ (e13 = op1(e11, e11)) | spl42_284), inference(avatar_component_clause, [], [f2763])).
fof(f2763, plain, (spl42_284 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_284])])).
fof(f15492, plain, ((e13 = op1(e11, e11)) | (~ spl42_267 | ~ spl42_331)), inference(forward_demodulation, [], [f15485, f2960])).
fof(f2960, plain, ((e11 = op1(e10, e11)) | ~ spl42_331), inference(avatar_component_clause, [], [f2958])).
fof(f2958, plain, (spl42_331 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_331])])).
fof(f15485, plain, ((e13 = op1(op1(e10, e11), e11)) | ~ spl42_267), inference(backward_demodulation, [], [f280, f2695])).
fof(f15477, plain, (~ spl42_349 | ~ spl42_350), inference(avatar_contradiction_clause, [], [f15476])).
fof(f15476, plain, ($false | (~ spl42_349 | ~ spl42_350)), inference(subsumption_resolution, [], [f15475, f1132])).
fof(f15475, plain, ((e25 = e26) | (~ spl42_349 | ~ spl42_350)), inference(forward_demodulation, [], [f3136, f3132])).
fof(f3136, plain, ((e26 = op2(e26, e26)) | ~ spl42_350), inference(avatar_component_clause, [], [f3134])).
fof(f3134, plain, (spl42_350 <=> (e26 = op2(e26, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_350])])).
fof(f15464, plain, (~ spl42_360 | ~ spl42_374), inference(avatar_split_clause, [], [f15454, f3234, f3176])).
fof(f15454, plain, (~ (e22 = op2(e26, e24)) | ~ spl42_374), inference(backward_demodulation, [], [f1082, f3236])).
fof(f15463, plain, (~ spl42_374 | ~ spl42_540 | spl42_640), inference(avatar_contradiction_clause, [], [f15462])).
fof(f15462, plain, ($false | (~ spl42_374 | ~ spl42_540 | spl42_640)), inference(subsumption_resolution, [], [f15461, f4337])).
fof(f4337, plain, (~ (e22 = op2(e20, e26)) | spl42_640), inference(avatar_component_clause, [], [f4336])).
fof(f4336, plain, (spl42_640 <=> (e22 = op2(e20, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_640])])).
fof(f15461, plain, ((e22 = op2(e20, e26)) | (~ spl42_374 | ~ spl42_540)), inference(forward_demodulation, [], [f15453, f3924])).
fof(f15453, plain, ((e22 = op2(op2(e22, e26), e26)) | ~ spl42_374), inference(backward_demodulation, [], [f474, f3236])).
fof(f15409, plain, (spl42_647 | ~ spl42_351 | ~ spl42_427), inference(avatar_split_clause, [], [f15408, f3453, f3139, f4365])).
fof(f4365, plain, (spl42_647 <=> (e22 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_647])])).
fof(f3453, plain, (spl42_427 <=> (e26 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_427])])).
fof(f15408, plain, ((e22 = op2(e20, e25)) | (~ spl42_351 | ~ spl42_427)), inference(forward_demodulation, [], [f15402, f3141])).
fof(f15402, plain, ((e22 = op2(op2(e26, e25), e25)) | ~ spl42_427), inference(backward_demodulation, [], [f473, f3455])).
fof(f3455, plain, ((e26 = op2(e25, e22)) | ~ spl42_427), inference(avatar_component_clause, [], [f3453])).
fof(f473, plain, (e22 = op2(op2(op2(e25, e22), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f15332, plain, (~ spl42_374 | ~ spl42_521), inference(avatar_split_clause, [], [f15325, f3843, f3234])).
fof(f3843, plain, (spl42_521 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_521])])).
fof(f15325, plain, (~ (e22 = op2(e26, e22)) | ~ spl42_521), inference(backward_demodulation, [], [f856, f3845])).
fof(f3845, plain, ((e22 = op2(e23, e22)) | ~ spl42_521), inference(avatar_component_clause, [], [f3843])).
fof(f856, plain, ~ (op2(e23, e22) = op2(e26, e22)), inference(cnf_transformation, [], [f8])).
fof(f15281, plain, (spl42_442 | ~ spl42_387 | ~ spl42_593), inference(avatar_split_clause, [], [f15276, f4141, f3288, f3516])).
fof(f4141, plain, (spl42_593 <=> (e24 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_593])])).
fof(f15276, plain, ((e20 = op2(e24, e26)) | (~ spl42_387 | ~ spl42_593)), inference(backward_demodulation, [], [f14305, f4143])).
fof(f4143, plain, ((e24 = op2(e21, e26)) | ~ spl42_593), inference(avatar_component_clause, [], [f4141])).
fof(f14305, plain, ((e20 = op2(op2(e21, e26), e26)) | ~ spl42_387), inference(backward_demodulation, [], [f460, f3290])).
fof(f15260, plain, (~ spl42_627 | spl42_866), inference(avatar_contradiction_clause, [], [f15259])).
fof(f15259, plain, ($false | (~ spl42_627 | spl42_866)), inference(subsumption_resolution, [], [f15258, f5733])).
fof(f5733, plain, (~ (e23 = h2(e15)) | spl42_866), inference(avatar_component_clause, [], [f5731])).
fof(f5731, plain, (spl42_866 <=> (e23 = h2(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_866])])).
fof(f15258, plain, ((e23 = h2(e15)) | ~ spl42_627), inference(backward_demodulation, [], [f1211, f4284])).
fof(f4284, plain, ((e23 = op2(e21, e21)) | ~ spl42_627), inference(avatar_component_clause, [], [f4282])).
fof(f15218, plain, (~ spl42_361 | ~ spl42_698), inference(avatar_split_clause, [], [f15205, f4872, f3180])).
fof(f15205, plain, (~ (e23 = op2(e26, e24)) | ~ spl42_698), inference(backward_demodulation, [], [f4799, f4873])).
fof(f4799, plain, ~ (op2(e26, e24) = h7(e13)), inference(backward_demodulation, [], [f1078, f4789])).
fof(f1078, plain, ~ (op2(e26, e21) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f15217, plain, (~ spl42_529 | ~ spl42_698), inference(avatar_split_clause, [], [f15203, f4872, f3876])).
fof(f15203, plain, (~ (e23 = op2(e23, e21)) | ~ spl42_698), inference(backward_demodulation, [], [f4793, f4873])).
fof(f4793, plain, ~ (op2(e23, e21) = h7(e13)), inference(backward_demodulation, [], [f835, f4789])).
fof(f835, plain, ~ (op2(e23, e21) = op2(e26, e21)), inference(cnf_transformation, [], [f8])).
fof(f15216, plain, (~ spl42_578 | ~ spl42_698), inference(avatar_split_clause, [], [f15202, f4872, f4079])).
fof(f15202, plain, (~ (e23 = op2(e22, e21)) | ~ spl42_698), inference(backward_demodulation, [], [f4792, f4873])).
fof(f4792, plain, ~ (op2(e22, e21) = h7(e13)), inference(backward_demodulation, [], [f832, f4789])).
fof(f832, plain, ~ (op2(e22, e21) = op2(e26, e21)), inference(cnf_transformation, [], [f8])).
fof(f15199, plain, (~ spl42_647 | ~ spl42_598), inference(avatar_split_clause, [], [f15198, f4162, f4365])).
fof(f15198, plain, (~ (e22 = op2(e20, e25)) | ~ spl42_598), inference(forward_demodulation, [], [f902, f4164])).
fof(f902, plain, ~ (op2(e20, e25) = op2(e21, e25)), inference(cnf_transformation, [], [f8])).
fof(f15103, plain, (~ spl42_608 | ~ spl42_510), inference(avatar_split_clause, [], [f15102, f3797, f4203])).
fof(f15102, plain, (~ (e25 = op2(e21, e24)) | ~ spl42_510), inference(forward_demodulation, [], [f888, f3799])).
fof(f3799, plain, ((e25 = op2(e23, e24)) | ~ spl42_510), inference(avatar_component_clause, [], [f3797])).
fof(f888, plain, ~ (op2(e21, e24) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f15043, plain, (~ spl42_482 | ~ spl42_468), inference(avatar_split_clause, [], [f15042, f3623, f3681])).
fof(f15042, plain, (~ (e25 = op2(e24, e21)) | ~ spl42_468), inference(forward_demodulation, [], [f1035, f3625])).
fof(f1035, plain, ~ (op2(e24, e21) = op2(e24, e23)), inference(cnf_transformation, [], [f8])).
fof(f15010, plain, (~ spl42_417 | ~ spl42_515), inference(avatar_split_clause, [], [f15004, f3818, f3412])).
fof(f15004, plain, (~ (e23 = op2(e25, e23)) | ~ spl42_515), inference(backward_demodulation, [], [f4712, f14194])).
fof(f14194, plain, ((e23 = h4(e15)) | ~ spl42_515), inference(backward_demodulation, [], [f1225, f3820])).
fof(f3820, plain, ((e23 = op2(e23, e23)) | ~ spl42_515), inference(avatar_component_clause, [], [f3818])).
fof(f4712, plain, ~ (op2(e25, e23) = h4(e15)), inference(backward_demodulation, [], [f876, f1225])).
fof(f876, plain, ~ (op2(e23, e23) = op2(e25, e23)), inference(cnf_transformation, [], [f8])).
fof(f14983, plain, (~ spl42_401 | ~ spl42_440 | spl42_596), inference(avatar_contradiction_clause, [], [f14982])).
fof(f14982, plain, ($false | (~ spl42_401 | ~ spl42_440 | spl42_596)), inference(subsumption_resolution, [], [f14981, f4155])).
fof(f4155, plain, (~ (e20 = op2(e21, e25)) | spl42_596), inference(avatar_component_clause, [], [f4154])).
fof(f4154, plain, (spl42_596 <=> (e20 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_596])])).
fof(f14981, plain, ((e20 = op2(e21, e25)) | (~ spl42_401 | ~ spl42_440)), inference(forward_demodulation, [], [f14980, f3348])).
fof(f14980, plain, ((e20 = op2(op2(e25, e25), e25)) | ~ spl42_440), inference(forward_demodulation, [], [f459, f3509])).
fof(f3509, plain, ((e25 = op2(e25, e20)) | ~ spl42_440), inference(avatar_component_clause, [], [f3507])).
fof(f3507, plain, (spl42_440 <=> (e25 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_440])])).
fof(f14965, plain, (~ spl42_398 | ~ spl42_688), inference(avatar_contradiction_clause, [], [f14964])).
fof(f14964, plain, ($false | (~ spl42_398 | ~ spl42_688)), inference(subsumption_resolution, [], [f14963, f1132])).
fof(f14963, plain, ((e25 = e26) | (~ spl42_398 | ~ spl42_688)), inference(forward_demodulation, [], [f14950, f3335])).
fof(f3335, plain, ((e25 = op2(e25, e26)) | ~ spl42_398), inference(avatar_component_clause, [], [f3333])).
fof(f3333, plain, (spl42_398 <=> (e25 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_398])])).
fof(f14950, plain, ((e26 = op2(e25, e26)) | (~ spl42_398 | ~ spl42_688)), inference(backward_demodulation, [], [f12687, f3335])).
fof(f14934, plain, (~ spl42_374 | ~ spl42_668), inference(avatar_split_clause, [], [f14933, f4452, f3234])).
fof(f4452, plain, (spl42_668 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_668])])).
fof(f14933, plain, (~ (e22 = op2(e26, e22)) | ~ spl42_668), inference(forward_demodulation, [], [f844, f4454])).
fof(f4454, plain, ((e22 = op2(e20, e22)) | ~ spl42_668), inference(avatar_component_clause, [], [f4452])).
fof(f844, plain, ~ (op2(e20, e22) = op2(e26, e22)), inference(cnf_transformation, [], [f8])).
fof(f14891, plain, (~ spl42_167 | ~ spl42_265), inference(avatar_contradiction_clause, [], [f14890])).
fof(f14890, plain, ($false | (~ spl42_167 | ~ spl42_265)), inference(subsumption_resolution, [], [f14889, f2686])).
fof(f14889, plain, (~ (e15 = op1(e11, e14)) | ~ spl42_167), inference(forward_demodulation, [], [f594, f2280])).
fof(f2280, plain, ((e15 = op1(e13, e14)) | ~ spl42_167), inference(avatar_component_clause, [], [f2278])).
fof(f2278, plain, (spl42_167 <=> (e15 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_167])])).
fof(f594, plain, ~ (op1(e11, e14) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f14837, plain, (~ spl42_26 | spl42_123 | ~ spl42_175), inference(avatar_contradiction_clause, [], [f14836])).
fof(f14836, plain, ($false | (~ spl42_26 | spl42_123 | ~ spl42_175)), inference(subsumption_resolution, [], [f14835, f2097])).
fof(f2097, plain, (~ (e13 = op1(e14, e13)) | spl42_123), inference(avatar_component_clause, [], [f2096])).
fof(f14835, plain, ((e13 = op1(e14, e13)) | (~ spl42_26 | ~ spl42_175)), inference(forward_demodulation, [], [f14834, f1696])).
fof(f14834, plain, ((e13 = op1(op1(e16, e13), e13)) | ~ spl42_175), inference(forward_demodulation, [], [f282, f2313])).
fof(f2313, plain, ((e16 = op1(e13, e13)) | ~ spl42_175), inference(avatar_component_clause, [], [f2311])).
fof(f2311, plain, (spl42_175 <=> (e16 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_175])])).
fof(f14803, plain, (~ spl42_157 | ~ spl42_255), inference(avatar_split_clause, [], [f13740, f2643, f2237])).
fof(f2237, plain, (spl42_157 <=> (e12 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_157])])).
fof(f13740, plain, (~ (e12 = op1(e13, e15)) | ~ spl42_255), inference(forward_demodulation, [], [f615, f2645])).
fof(f615, plain, ~ (op1(e11, e15) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f14761, plain, (spl42_105 | ~ spl42_6 | ~ spl42_54), inference(avatar_split_clause, [], [f14623, f1810, f1611, f2021])).
fof(f1810, plain, (spl42_54 <=> (e14 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_54])])).
fof(f14623, plain, ((e16 = op1(e14, e16)) | (~ spl42_6 | ~ spl42_54)), inference(backward_demodulation, [], [f12249, f1812])).
fof(f1812, plain, ((e14 = op1(e15, e16)) | ~ spl42_54), inference(avatar_component_clause, [], [f1810])).
fof(f14742, plain, (~ spl42_132 | ~ spl42_83), inference(avatar_split_clause, [], [f14733, f1930, f2133])).
fof(f14733, plain, (~ (e15 = op1(e14, e12)) | ~ spl42_83), inference(backward_demodulation, [], [f563, f1932])).
fof(f563, plain, ~ (op1(e14, e12) = op1(e15, e12)), inference(cnf_transformation, [], [f7])).
fof(f14725, plain, (~ spl42_74 | ~ spl42_319), inference(avatar_split_clause, [], [f14352, f2908, f1893])).
fof(f2908, plain, (spl42_319 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_319])])).
fof(f14352, plain, (~ (e13 = op1(e15, e13)) | ~ spl42_319), inference(backward_demodulation, [], [f570, f2910])).
fof(f2910, plain, ((e13 = op1(e10, e13)) | ~ spl42_319), inference(avatar_component_clause, [], [f2908])).
fof(f570, plain, ~ (op1(e10, e13) = op1(e15, e13)), inference(cnf_transformation, [], [f7])).
fof(f14702, plain, (~ spl42_58 | ~ spl42_66 | ~ spl42_209), inference(avatar_contradiction_clause, [], [f14701])).
fof(f14701, plain, ($false | (~ spl42_58 | ~ spl42_66 | ~ spl42_209)), inference(subsumption_resolution, [], [f14700, f1099])).
fof(f14700, plain, ((e11 = e14) | (~ spl42_58 | ~ spl42_66 | ~ spl42_209)), inference(forward_demodulation, [], [f14699, f1829])).
fof(f14699, plain, ((e14 = op1(e15, e15)) | (~ spl42_66 | ~ spl42_209)), inference(forward_demodulation, [], [f14698, f2454])).
fof(f14698, plain, ((e14 = op1(op1(e12, e15), e15)) | ~ spl42_66), inference(forward_demodulation, [], [f291, f1862])).
fof(f1862, plain, ((e12 = op1(e15, e14)) | ~ spl42_66), inference(avatar_component_clause, [], [f1860])).
fof(f1860, plain, (spl42_66 <=> (e12 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_66])])).
fof(f14690, plain, (~ spl42_46 | ~ spl42_39), inference(avatar_split_clause, [], [f12260, f1748, f1777])).
fof(f1777, plain, (spl42_46 <=> (e13 = op1(e16, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_46])])).
fof(f12260, plain, (~ (e13 = op1(e16, e10)) | ~ spl42_39), inference(forward_demodulation, [], [f776, f1750])).
fof(f776, plain, ~ (op1(e16, e10) = op1(e16, e11)), inference(cnf_transformation, [], [f7])).
fof(f14685, plain, (~ spl42_18 | ~ spl42_39), inference(avatar_split_clause, [], [f13569, f1748, f1661])).
fof(f1661, plain, (spl42_18 <=> (e13 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_18])])).
fof(f13569, plain, (~ (e13 = op1(e16, e14)) | ~ spl42_39), inference(backward_demodulation, [], [f784, f1750])).
fof(f784, plain, ~ (op1(e16, e11) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f14683, plain, (~ spl42_6 | ~ spl42_7), inference(avatar_contradiction_clause, [], [f14682])).
fof(f14682, plain, ($false | (~ spl42_6 | ~ spl42_7)), inference(subsumption_resolution, [], [f14681, f1111])).
fof(f14681, plain, ((e15 = e16) | (~ spl42_6 | ~ spl42_7)), inference(forward_demodulation, [], [f1617, f1613])).
fof(f1617, plain, ((e16 = op1(e16, e16)) | ~ spl42_7), inference(avatar_component_clause, [], [f1615])).
fof(f1615, plain, (spl42_7 <=> (e16 = op1(e16, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_7])])).
fof(f14611, plain, (~ spl42_58 | ~ spl42_73 | ~ spl42_209), inference(avatar_contradiction_clause, [], [f14610])).
fof(f14610, plain, ($false | (~ spl42_58 | ~ spl42_73 | ~ spl42_209)), inference(subsumption_resolution, [], [f14609, f1098])).
fof(f14609, plain, ((e11 = e13) | (~ spl42_58 | ~ spl42_73 | ~ spl42_209)), inference(forward_demodulation, [], [f14608, f1829])).
fof(f14608, plain, ((e13 = op1(e15, e15)) | (~ spl42_73 | ~ spl42_209)), inference(forward_demodulation, [], [f14601, f2454])).
fof(f14601, plain, ((e13 = op1(op1(e12, e15), e15)) | ~ spl42_73), inference(backward_demodulation, [], [f284, f1891])).
fof(f1891, plain, ((e12 = op1(e15, e13)) | ~ spl42_73), inference(avatar_component_clause, [], [f1889])).
fof(f284, plain, (e13 = op1(op1(op1(e15, e13), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f14576, plain, (~ spl42_96 | ~ spl42_109 | spl42_155), inference(avatar_contradiction_clause, [], [f14575])).
fof(f14575, plain, ($false | (~ spl42_96 | ~ spl42_109 | spl42_155)), inference(subsumption_resolution, [], [f14570, f2230])).
fof(f2230, plain, (~ (e10 = op1(e13, e15)) | spl42_155), inference(avatar_component_clause, [], [f2229])).
fof(f2229, plain, (spl42_155 <=> (e10 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_155])])).
fof(f14570, plain, ((e10 = op1(e13, e15)) | (~ spl42_96 | ~ spl42_109)), inference(backward_demodulation, [], [f13655, f2040])).
fof(f2040, plain, ((e13 = op1(e14, e15)) | ~ spl42_109), inference(avatar_component_clause, [], [f2038])).
fof(f2038, plain, (spl42_109 <=> (e13 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_109])])).
fof(f13655, plain, ((e10 = op1(op1(e14, e15), e15)) | ~ spl42_96), inference(forward_demodulation, [], [f263, f1986])).
fof(f14549, plain, (spl42_30 | ~ spl42_133 | ~ spl42_236), inference(avatar_contradiction_clause, [], [f14548])).
fof(f14548, plain, ($false | (spl42_30 | ~ spl42_133 | ~ spl42_236)), inference(subsumption_resolution, [], [f14544, f1712])).
fof(f1712, plain, (~ (e11 = op1(e16, e12)) | spl42_30), inference(avatar_component_clause, [], [f1711])).
fof(f14544, plain, ((e11 = op1(e16, e12)) | (~ spl42_133 | ~ spl42_236)), inference(backward_demodulation, [], [f14421, f2139])).
fof(f14421, plain, ((e11 = op1(op1(e14, e12), e12)) | ~ spl42_236), inference(backward_demodulation, [], [f267, f2566])).
fof(f267, plain, (e11 = op1(op1(op1(e12, e11), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f14520, plain, (~ spl42_39 | ~ spl42_150 | spl42_198), inference(avatar_contradiction_clause, [], [f14519])).
fof(f14519, plain, ($false | (~ spl42_39 | ~ spl42_150 | spl42_198)), inference(subsumption_resolution, [], [f14513, f2408])).
fof(f2408, plain, (~ (e11 = op1(e12, e16)) | spl42_198), inference(avatar_component_clause, [], [f2407])).
fof(f2407, plain, (spl42_198 <=> (e11 = op1(e12, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_198])])).
fof(f14513, plain, ((e11 = op1(e12, e16)) | (~ spl42_39 | ~ spl42_150)), inference(backward_demodulation, [], [f12233, f2210])).
fof(f14518, plain, (~ spl42_101 | ~ spl42_150), inference(avatar_split_clause, [], [f14507, f2208, f2005])).
fof(f2005, plain, (spl42_101 <=> (e12 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_101])])).
fof(f14507, plain, (~ (e12 = op1(e14, e16)) | ~ spl42_150), inference(backward_demodulation, [], [f644, f2210])).
fof(f644, plain, ~ (op1(e13, e16) = op1(e14, e16)), inference(cnf_transformation, [], [f7])).
fof(f14469, plain, (~ spl42_158 | ~ spl42_193), inference(avatar_split_clause, [], [f14460, f2386, f2241])).
fof(f2386, plain, (spl42_193 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_193])])).
fof(f14460, plain, (~ (e13 = op1(e13, e15)) | ~ spl42_193), inference(backward_demodulation, [], [f717, f2388])).
fof(f2388, plain, ((e13 = op1(e13, e10)) | ~ spl42_193), inference(avatar_component_clause, [], [f2386])).
fof(f717, plain, ~ (op1(e13, e10) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f14427, plain, (~ spl42_229 | ~ spl42_236), inference(avatar_split_clause, [], [f14423, f2564, f2535])).
fof(f14423, plain, (~ (e14 = op1(e12, e12)) | ~ spl42_236), inference(backward_demodulation, [], [f698, f2566])).
fof(f698, plain, ~ (op1(e12, e11) = op1(e12, e12)), inference(cnf_transformation, [], [f7])).
fof(f14304, plain, (~ spl42_401 | ~ spl42_406), inference(avatar_contradiction_clause, [], [f14303])).
fof(f14303, plain, ($false | (~ spl42_401 | ~ spl42_406)), inference(subsumption_resolution, [], [f14302, f1122])).
fof(f14302, plain, ((e21 = e26) | (~ spl42_401 | ~ spl42_406)), inference(forward_demodulation, [], [f3368, f3348])).
fof(f3368, plain, ((e26 = op2(e25, e25)) | ~ spl42_406), inference(avatar_component_clause, [], [f3366])).
fof(f3366, plain, (spl42_406 <=> (e26 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_406])])).
fof(f14211, plain, (~ spl42_444 | ~ spl42_493), inference(avatar_split_clause, [], [f14201, f3727, f3524])).
fof(f14201, plain, (~ (e22 = op2(e24, e26)) | ~ spl42_493), inference(backward_demodulation, [], [f938, f3729])).
fof(f938, plain, ~ (op2(e23, e26) = op2(e24, e26)), inference(cnf_transformation, [], [f8])).
fof(f14136, plain, (spl42_831 | ~ spl42_571), inference(avatar_split_clause, [], [f14135, f4050, f5555])).
fof(f14135, plain, ((e23 = h3(e15)) | ~ spl42_571), inference(backward_demodulation, [], [f1218, f4052])).
fof(f14086, plain, (spl42_873 | ~ spl42_626), inference(avatar_split_clause, [], [f14085, f4278, f5766])).
fof(f5766, plain, (spl42_873 <=> (e22 = h2(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_873])])).
fof(f4278, plain, (spl42_626 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_626])])).
fof(f14085, plain, ((e22 = h2(e15)) | ~ spl42_626), inference(backward_demodulation, [], [f1211, f4280])).
fof(f4280, plain, ((e22 = op2(e21, e21)) | ~ spl42_626), inference(avatar_component_clause, [], [f4278])).
fof(f14084, plain, (~ spl42_642 | ~ spl42_644), inference(avatar_contradiction_clause, [], [f14083])).
fof(f14083, plain, ($false | (~ spl42_642 | ~ spl42_644)), inference(subsumption_resolution, [], [f14082, f1131])).
fof(f14082, plain, ((e24 = e26) | (~ spl42_642 | ~ spl42_644)), inference(backward_demodulation, [], [f4354, f4346])).
fof(f4346, plain, ((e24 = op2(e20, e26)) | ~ spl42_642), inference(avatar_component_clause, [], [f4344])).
fof(f4344, plain, (spl42_642 <=> (e24 = op2(e20, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_642])])).
fof(f14044, plain, (~ spl42_866 | ~ spl42_382), inference(avatar_split_clause, [], [f14043, f3267, f5731])).
fof(f14043, plain, (~ (e23 = h2(e15)) | ~ spl42_382), inference(forward_demodulation, [], [f4671, f3269])).
fof(f4671, plain, ~ (op2(e26, e21) = h2(e15)), inference(backward_demodulation, [], [f828, f1211])).
fof(f828, plain, ~ (op2(e21, e21) = op2(e26, e21)), inference(cnf_transformation, [], [f8])).
fof(f14028, plain, (~ spl42_873 | ~ spl42_598), inference(avatar_split_clause, [], [f14027, f4162, f5766])).
fof(f14027, plain, (~ (e22 = h2(e15)) | ~ spl42_598), inference(forward_demodulation, [], [f4676, f4164])).
fof(f4676, plain, ~ (op2(e21, e25) = h2(e15)), inference(backward_demodulation, [], [f974, f1211])).
fof(f974, plain, ~ (op2(e21, e21) = op2(e21, e25)), inference(cnf_transformation, [], [f8])).
fof(f14002, plain, (spl42_552 | ~ spl42_598 | ~ spl42_733), inference(avatar_contradiction_clause, [], [f14001])).
fof(f14001, plain, ($false | (spl42_552 | ~ spl42_598 | ~ spl42_733)), inference(subsumption_resolution, [], [f13007, f3972])).
fof(f3972, plain, (~ (e25 = op2(e22, e25)) | spl42_552), inference(avatar_component_clause, [], [f3971])).
fof(f13007, plain, ((e25 = op2(e22, e25)) | (~ spl42_598 | ~ spl42_733)), inference(forward_demodulation, [], [f13001, f4164])).
fof(f13001, plain, ((e25 = op2(op2(e21, e25), e25)) | ~ spl42_733), inference(backward_demodulation, [], [f4749, f5064])).
fof(f4749, plain, (e25 = op2(op2(h6(e15), e25), e25)), inference(backward_demodulation, [], [f494, f1239])).
fof(f494, plain, (e25 = op2(op2(op2(e25, e25), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f13900, plain, (~ spl42_488 | ~ spl42_439), inference(avatar_split_clause, [], [f13899, f3503, f3706])).
fof(f13899, plain, (~ (e24 = op2(e24, e20)) | ~ spl42_439), inference(forward_demodulation, [], [f815, f3505])).
fof(f815, plain, ~ (op2(e24, e20) = op2(e25, e20)), inference(cnf_transformation, [], [f8])).
fof(f13858, plain, (spl42_497 | ~ spl42_396 | ~ spl42_688), inference(avatar_split_clause, [], [f13857, f4816, f3325, f3743])).
fof(f3325, plain, (spl42_396 <=> (e23 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_396])])).
fof(f13857, plain, ((e26 = op2(e23, e26)) | (~ spl42_396 | ~ spl42_688)), inference(forward_demodulation, [], [f12687, f3327])).
fof(f3327, plain, ((e23 = op2(e25, e26)) | ~ spl42_396), inference(avatar_component_clause, [], [f3325])).
fof(f13826, plain, (spl42_698 | ~ spl42_382), inference(avatar_split_clause, [], [f13825, f3267, f4872])).
fof(f13825, plain, ((e23 = h7(e13)) | ~ spl42_382), inference(forward_demodulation, [], [f4789, f3269])).
fof(f13798, plain, (~ spl42_358 | ~ spl42_351), inference(avatar_split_clause, [], [f12582, f3139, f3168])).
fof(f3168, plain, (spl42_358 <=> (e20 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_358])])).
fof(f12582, plain, (~ (e20 = op2(e26, e24)) | ~ spl42_351), inference(forward_demodulation, [], [f1088, f3141])).
fof(f1088, plain, ~ (op2(e26, e24) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f13797, plain, (~ spl42_345 | ~ spl42_349), inference(avatar_contradiction_clause, [], [f13796])).
fof(f13796, plain, ($false | (~ spl42_345 | ~ spl42_349)), inference(subsumption_resolution, [], [f13795, f1121])).
fof(f13795, plain, ((e21 = e25) | (~ spl42_345 | ~ spl42_349)), inference(forward_demodulation, [], [f3132, f3116])).
fof(f3116, plain, ((e21 = op2(e26, e26)) | ~ spl42_345), inference(avatar_component_clause, [], [f3114])).
fof(f13786, plain, (~ spl42_304 | ~ spl42_255), inference(avatar_split_clause, [], [f13785, f2643, f2846])).
fof(f2846, plain, (spl42_304 <=> (e12 = op1(e10, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_304])])).
fof(f13785, plain, (~ (e12 = op1(e10, e15)) | ~ spl42_255), inference(forward_demodulation, [], [f608, f2645])).
fof(f608, plain, ~ (op1(e10, e15) = op1(e11, e15)), inference(cnf_transformation, [], [f7])).
fof(f13750, plain, (~ spl42_285 | ~ spl42_264), inference(avatar_split_clause, [], [f13749, f2680, f2767])).
fof(f2767, plain, (spl42_285 <=> (e14 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_285])])).
fof(f2680, plain, (spl42_264 <=> (e14 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_264])])).
fof(f13749, plain, (~ (e14 = op1(e11, e11)) | ~ spl42_264), inference(forward_demodulation, [], [f679, f2682])).
fof(f2682, plain, ((e14 = op1(e11, e14)) | ~ spl42_264), inference(avatar_component_clause, [], [f2680])).
fof(f679, plain, ~ (op1(e11, e11) = op1(e11, e14)), inference(cnf_transformation, [], [f7])).
fof(f13748, plain, (~ spl42_278 | ~ spl42_264), inference(avatar_split_clause, [], [f13747, f2680, f2738])).
fof(f13747, plain, (~ (e14 = op1(e11, e12)) | ~ spl42_264), inference(forward_demodulation, [], [f683, f2682])).
fof(f683, plain, ~ (op1(e11, e12) = op1(e11, e14)), inference(cnf_transformation, [], [f7])).
fof(f13642, plain, (~ spl42_58 | ~ spl42_90 | spl42_254), inference(avatar_contradiction_clause, [], [f13641])).
fof(f13641, plain, ($false | (~ spl42_58 | ~ spl42_90 | spl42_254)), inference(subsumption_resolution, [], [f13640, f2640])).
fof(f2640, plain, (~ (e11 = op1(e11, e15)) | spl42_254), inference(avatar_component_clause, [], [f2639])).
fof(f13640, plain, ((e11 = op1(e11, e15)) | (~ spl42_58 | ~ spl42_90)), inference(forward_demodulation, [], [f13639, f1829])).
fof(f13639, plain, ((e11 = op1(op1(e15, e15), e15)) | ~ spl42_90), inference(forward_demodulation, [], [f270, f1961])).
fof(f1961, plain, ((e15 = op1(e15, e11)) | ~ spl42_90), inference(avatar_component_clause, [], [f1959])).
fof(f270, plain, (e11 = op1(op1(op1(e15, e11), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f13537, plain, (~ spl42_36 | ~ spl42_39), inference(avatar_contradiction_clause, [], [f13536])).
fof(f13536, plain, ($false | (~ spl42_36 | ~ spl42_39)), inference(subsumption_resolution, [], [f13535, f1093])).
fof(f1093, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f9])).
fof(f13535, plain, ((e10 = e13) | (~ spl42_36 | ~ spl42_39)), inference(backward_demodulation, [], [f1750, f1738])).
fof(f1738, plain, ((e10 = op1(e16, e11)) | ~ spl42_36), inference(avatar_component_clause, [], [f1736])).
fof(f1736, plain, (spl42_36 <=> (e10 = op1(e16, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_36])])).
fof(f13531, plain, (~ spl42_39 | ~ spl42_42), inference(avatar_contradiction_clause, [], [f13530])).
fof(f13530, plain, ($false | (~ spl42_39 | ~ spl42_42)), inference(subsumption_resolution, [], [f13529, f1108])).
fof(f13529, plain, ((e13 = e16) | (~ spl42_39 | ~ spl42_42)), inference(forward_demodulation, [], [f1762, f1750])).
fof(f1762, plain, ((e16 = op1(e16, e11)) | ~ spl42_42), inference(avatar_component_clause, [], [f1760])).
fof(f1760, plain, (spl42_42 <=> (e16 = op1(e16, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_42])])).
fof(f13493, plain, (~ spl42_80 | ~ spl42_83), inference(avatar_contradiction_clause, [], [f13492])).
fof(f13492, plain, ($false | (~ spl42_80 | ~ spl42_83)), inference(subsumption_resolution, [], [f13491, f1104])).
fof(f13491, plain, ((e12 = e15) | (~ spl42_80 | ~ spl42_83)), inference(forward_demodulation, [], [f1932, f1920])).
fof(f1920, plain, ((e12 = op1(e15, e12)) | ~ spl42_80), inference(avatar_component_clause, [], [f1918])).
fof(f1918, plain, (spl42_80 <=> (e12 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_80])])).
fof(f13457, plain, (spl42_263 | ~ spl42_16 | ~ spl42_126), inference(avatar_split_clause, [], [f13456, f2108, f1653, f2676])).
fof(f2108, plain, (spl42_126 <=> (e16 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_126])])).
fof(f13456, plain, ((e13 = op1(e11, e14)) | (~ spl42_16 | ~ spl42_126)), inference(forward_demodulation, [], [f13453, f1655])).
fof(f13453, plain, ((e13 = op1(op1(e16, e14), e14)) | ~ spl42_126), inference(backward_demodulation, [], [f283, f2110])).
fof(f2110, plain, ((e16 = op1(e14, e13)) | ~ spl42_126), inference(avatar_component_clause, [], [f2108])).
fof(f283, plain, (e13 = op1(op1(op1(e14, e13), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f13409, plain, (spl42_124 | ~ spl42_26 | ~ spl42_168), inference(avatar_split_clause, [], [f13408, f2282, f1694, f2100])).
fof(f2100, plain, (spl42_124 <=> (e14 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_124])])).
fof(f2282, plain, (spl42_168 <=> (e16 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_168])])).
fof(f13408, plain, ((e14 = op1(e14, e13)) | (~ spl42_26 | ~ spl42_168)), inference(forward_demodulation, [], [f13406, f1696])).
fof(f13406, plain, ((e14 = op1(op1(e16, e13), e13)) | ~ spl42_168), inference(backward_demodulation, [], [f289, f2284])).
fof(f2284, plain, ((e16 = op1(e13, e14)) | ~ spl42_168), inference(avatar_component_clause, [], [f2282])).
fof(f13343, plain, (~ spl42_215 | ~ spl42_264), inference(avatar_split_clause, [], [f13339, f2680, f2477])).
fof(f13339, plain, (~ (e14 = op1(e12, e14)) | ~ spl42_264), inference(backward_demodulation, [], [f593, f2682])).
fof(f593, plain, ~ (op1(e11, e14) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f13275, plain, (~ spl42_122 | ~ spl42_318), inference(avatar_split_clause, [], [f13268, f2904, f2092])).
fof(f2904, plain, (spl42_318 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_318])])).
fof(f13268, plain, (~ (e12 = op1(e14, e13)) | ~ spl42_318), inference(backward_demodulation, [], [f569, f2906])).
fof(f2906, plain, ((e12 = op1(e10, e13)) | ~ spl42_318), inference(avatar_component_clause, [], [f2904])).
fof(f569, plain, ~ (op1(e10, e13) = op1(e14, e13)), inference(cnf_transformation, [], [f7])).
fof(f13245, plain, (~ spl42_348 | ~ spl42_349), inference(avatar_contradiction_clause, [], [f13244])).
fof(f13244, plain, ($false | (~ spl42_348 | ~ spl42_349)), inference(subsumption_resolution, [], [f13243, f1130])).
fof(f13243, plain, ((e24 = e25) | (~ spl42_348 | ~ spl42_349)), inference(backward_demodulation, [], [f3132, f3128])).
fof(f3128, plain, ((e24 = op2(e26, e26)) | ~ spl42_348), inference(avatar_component_clause, [], [f3126])).
fof(f3126, plain, (spl42_348 <=> (e24 = op2(e26, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_348])])).
fof(f13215, plain, (~ spl42_425 | ~ spl42_426), inference(avatar_contradiction_clause, [], [f13214])).
fof(f13214, plain, ($false | (~ spl42_425 | ~ spl42_426)), inference(subsumption_resolution, [], [f13213, f1130])).
fof(f13213, plain, ((e24 = e25) | (~ spl42_425 | ~ spl42_426)), inference(forward_demodulation, [], [f3451, f3447])).
fof(f3447, plain, ((e24 = op2(e25, e22)) | ~ spl42_425), inference(avatar_component_clause, [], [f3445])).
fof(f3445, plain, (spl42_425 <=> (e24 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_425])])).
fof(f13176, plain, (~ spl42_453 | ~ spl42_481), inference(avatar_contradiction_clause, [], [f13175])).
fof(f13175, plain, ($false | (~ spl42_453 | ~ spl42_481)), inference(subsumption_resolution, [], [f13174, f1121])).
fof(f13174, plain, ((e21 = e25) | (~ spl42_453 | ~ spl42_481)), inference(backward_demodulation, [], [f12800, f13173])).
fof(f13173, plain, ((e21 = op2(h5(e15), e24)) | ~ spl42_481), inference(forward_demodulation, [], [f13168, f1232])).
fof(f13168, plain, ((e21 = op2(op2(e24, e24), e24)) | ~ spl42_481), inference(backward_demodulation, [], [f465, f3679])).
fof(f12800, plain, ((e25 = op2(h5(e15), e24)) | ~ spl42_453), inference(backward_demodulation, [], [f12799, f1232])).
fof(f12799, plain, ((e25 = op2(op2(e24, e24), e24)) | ~ spl42_453), inference(forward_demodulation, [], [f493, f3563])).
fof(f13050, plain, (~ spl42_556 | ~ spl42_654), inference(avatar_split_clause, [], [f13043, f4394, f3988])).
fof(f13043, plain, (~ (e22 = op2(e22, e24)) | ~ spl42_654), inference(backward_demodulation, [], [f882, f4396])).
fof(f882, plain, ~ (op2(e20, e24) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f13014, plain, (spl42_731 | ~ spl42_430 | ~ spl42_733), inference(avatar_split_clause, [], [f13013, f5063, f3466, f5053])).
fof(f13013, plain, ((e22 = h6(e10)) | (~ spl42_430 | ~ spl42_733)), inference(forward_demodulation, [], [f13004, f3468])).
fof(f13010, plain, (~ spl42_551 | ~ spl42_598 | ~ spl42_733), inference(avatar_contradiction_clause, [], [f13009])).
fof(f13009, plain, ($false | (~ spl42_551 | ~ spl42_598 | ~ spl42_733)), inference(subsumption_resolution, [], [f13008, f1130])).
fof(f13008, plain, ((e24 = e25) | (~ spl42_551 | ~ spl42_598 | ~ spl42_733)), inference(forward_demodulation, [], [f13007, f3969])).
fof(f3969, plain, ((e24 = op2(e22, e25)) | ~ spl42_551), inference(avatar_component_clause, [], [f3967])).
fof(f3967, plain, (spl42_551 <=> (e24 = op2(e22, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_551])])).
fof(f12963, plain, (~ spl42_619 | ~ spl42_598), inference(avatar_split_clause, [], [f12962, f4162, f4249])).
fof(f4249, plain, (spl42_619 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_619])])).
fof(f12962, plain, (~ (e22 = op2(e21, e22)) | ~ spl42_598), inference(forward_demodulation, [], [f978, f4164])).
fof(f978, plain, ~ (op2(e21, e22) = op2(e21, e25)), inference(cnf_transformation, [], [f8])).
fof(f12951, plain, (~ spl42_612 | ~ spl42_598), inference(avatar_split_clause, [], [f12950, f4162, f4220])).
fof(f4220, plain, (spl42_612 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_612])])).
fof(f12950, plain, (~ (e22 = op2(e21, e23)) | ~ spl42_598), inference(forward_demodulation, [], [f981, f4164])).
fof(f981, plain, ~ (op2(e21, e23) = op2(e21, e25)), inference(cnf_transformation, [], [f8])).
fof(f12949, plain, (~ spl42_611 | ~ spl42_625), inference(avatar_split_clause, [], [f12948, f4274, f4216])).
fof(f4274, plain, (spl42_625 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_625])])).
fof(f12948, plain, (~ (e21 = op2(e21, e23)) | ~ spl42_625), inference(forward_demodulation, [], [f4674, f11300])).
fof(f11300, plain, ((e21 = h2(e15)) | ~ spl42_625), inference(backward_demodulation, [], [f1211, f4276])).
fof(f4276, plain, ((e21 = op2(e21, e21)) | ~ spl42_625), inference(avatar_component_clause, [], [f4274])).
fof(f4674, plain, ~ (op2(e21, e23) = h2(e15)), inference(backward_demodulation, [], [f972, f1211])).
fof(f972, plain, ~ (op2(e21, e21) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f12947, plain, (~ spl42_607 | ~ spl42_411), inference(avatar_split_clause, [], [f12946, f3387, f4199])).
fof(f3387, plain, (spl42_411 <=> (e24 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_411])])).
fof(f12946, plain, (~ (e24 = op2(e21, e24)) | ~ spl42_411), inference(forward_demodulation, [], [f890, f3389])).
fof(f3389, plain, ((e24 = op2(e25, e24)) | ~ spl42_411), inference(avatar_component_clause, [], [f3387])).
fof(f890, plain, ~ (op2(e21, e24) = op2(e25, e24)), inference(cnf_transformation, [], [f8])).
fof(f12945, plain, (~ spl42_605 | ~ spl42_598), inference(avatar_split_clause, [], [f12944, f4162, f4191])).
fof(f4191, plain, (spl42_605 <=> (e22 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_605])])).
fof(f12944, plain, (~ (e22 = op2(e21, e24)) | ~ spl42_598), inference(forward_demodulation, [], [f983, f4164])).
fof(f983, plain, ~ (op2(e21, e24) = op2(e21, e25)), inference(cnf_transformation, [], [f8])).
fof(f12941, plain, (~ spl42_598 | ~ spl42_602), inference(avatar_contradiction_clause, [], [f12940])).
fof(f12940, plain, ($false | (~ spl42_598 | ~ spl42_602)), inference(subsumption_resolution, [], [f12939, f1126])).
fof(f1126, plain, ~ (e22 = e26), inference(cnf_transformation, [], [f10])).
fof(f12939, plain, ((e22 = e26) | (~ spl42_598 | ~ spl42_602)), inference(forward_demodulation, [], [f4180, f4164])).
fof(f4180, plain, ((e26 = op2(e21, e25)) | ~ spl42_602), inference(avatar_component_clause, [], [f4178])).
fof(f4178, plain, (spl42_602 <=> (e26 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_602])])).
fof(f12924, plain, (~ spl42_594 | ~ spl42_688), inference(avatar_split_clause, [], [f12923, f4816, f4145])).
fof(f12923, plain, (~ (e25 = op2(e21, e26)) | ~ spl42_688), inference(forward_demodulation, [], [f4772, f4817])).
fof(f4772, plain, ~ (op2(e21, e26) = h7(e15)), inference(backward_demodulation, [], [f933, f1246])).
fof(f933, plain, ~ (op2(e21, e26) = op2(e26, e26)), inference(cnf_transformation, [], [f8])).
fof(f12876, plain, (~ spl42_545 | ~ spl42_688), inference(avatar_split_clause, [], [f12875, f4816, f3942])).
fof(f3942, plain, (spl42_545 <=> (e25 = op2(e22, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_545])])).
fof(f12875, plain, (~ (e25 = op2(e22, e26)) | ~ spl42_688), inference(forward_demodulation, [], [f4773, f4817])).
fof(f4773, plain, ~ (op2(e22, e26) = h7(e15)), inference(backward_demodulation, [], [f937, f1246])).
fof(f937, plain, ~ (op2(e22, e26) = op2(e26, e26)), inference(cnf_transformation, [], [f8])).
fof(f12866, plain, (~ spl42_528 | ~ spl42_430), inference(avatar_split_clause, [], [f12865, f3466, f3872])).
fof(f12865, plain, (~ (e22 = op2(e23, e21)) | ~ spl42_430), inference(forward_demodulation, [], [f834, f3468])).
fof(f834, plain, ~ (op2(e23, e21) = op2(e25, e21)), inference(cnf_transformation, [], [f8])).
fof(f12848, plain, (~ spl42_438 | spl42_449 | ~ spl42_502), inference(avatar_contradiction_clause, [], [f12847])).
fof(f12847, plain, ($false | (~ spl42_438 | spl42_449 | ~ spl42_502)), inference(subsumption_resolution, [], [f12846, f3546])).
fof(f12846, plain, ((e20 = op2(e24, e25)) | (~ spl42_438 | ~ spl42_502)), inference(backward_demodulation, [], [f12763, f3766])).
fof(f12763, plain, ((e20 = op2(op2(e23, e25), e25)) | ~ spl42_438), inference(forward_demodulation, [], [f459, f3501])).
fof(f3501, plain, ((e23 = op2(e25, e20)) | ~ spl42_438), inference(avatar_component_clause, [], [f3499])).
fof(f3499, plain, (spl42_438 <=> (e23 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_438])])).
fof(f12819, plain, (~ spl42_496 | ~ spl42_688), inference(avatar_split_clause, [], [f12818, f4816, f3739])).
fof(f3739, plain, (spl42_496 <=> (e25 = op2(e23, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_496])])).
fof(f12818, plain, (~ (e25 = op2(e23, e26)) | ~ spl42_688), inference(forward_demodulation, [], [f4774, f4817])).
fof(f4774, plain, ~ (op2(e23, e26) = h7(e15)), inference(backward_demodulation, [], [f940, f1246])).
fof(f940, plain, ~ (op2(e23, e26) = op2(e26, e26)), inference(cnf_transformation, [], [f8])).
fof(f12765, plain, (~ spl42_447 | ~ spl42_688), inference(avatar_split_clause, [], [f12764, f4816, f3536])).
fof(f3536, plain, (spl42_447 <=> (e25 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_447])])).
fof(f12764, plain, (~ (e25 = op2(e24, e26)) | ~ spl42_688), inference(forward_demodulation, [], [f4775, f4817])).
fof(f4775, plain, ~ (op2(e24, e26) = h7(e15)), inference(backward_demodulation, [], [f942, f1246])).
fof(f942, plain, ~ (op2(e24, e26) = op2(e26, e26)), inference(cnf_transformation, [], [f8])).
fof(f12734, plain, (~ spl42_401 | ~ spl42_419 | ~ spl42_598), inference(avatar_contradiction_clause, [], [f12733])).
fof(f12733, plain, ($false | (~ spl42_401 | ~ spl42_419 | ~ spl42_598)), inference(subsumption_resolution, [], [f12732, f1123])).
fof(f1123, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f10])).
fof(f12732, plain, ((e22 = e23) | (~ spl42_401 | ~ spl42_419 | ~ spl42_598)), inference(forward_demodulation, [], [f12731, f4164])).
fof(f12731, plain, ((e23 = op2(e21, e25)) | (~ spl42_401 | ~ spl42_419)), inference(forward_demodulation, [], [f12730, f3348])).
fof(f12730, plain, ((e23 = op2(op2(e25, e25), e25)) | ~ spl42_419), inference(forward_demodulation, [], [f480, f3422])).
fof(f3422, plain, ((e25 = op2(e25, e23)) | ~ spl42_419), inference(avatar_component_clause, [], [f3420])).
fof(f3420, plain, (spl42_419 <=> (e25 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_419])])).
fof(f12695, plain, (~ spl42_390 | ~ spl42_369), inference(avatar_split_clause, [], [f12694, f3213, f3300])).
fof(f3300, plain, (spl42_390 <=> (e24 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_390])])).
fof(f12694, plain, (~ (e24 = op2(e26, e20)) | ~ spl42_369), inference(forward_demodulation, [], [f1072, f3215])).
fof(f1072, plain, ~ (op2(e26, e20) = op2(e26, e23)), inference(cnf_transformation, [], [f8])).
fof(f12691, plain, (~ spl42_391 | ~ spl42_688), inference(avatar_split_clause, [], [f12690, f4816, f3304])).
fof(f3304, plain, (spl42_391 <=> (e25 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_391])])).
fof(f12690, plain, (~ (e25 = op2(e26, e20)) | ~ spl42_688), inference(forward_demodulation, [], [f4777, f4817])).
fof(f4777, plain, ~ (op2(e26, e20) = h7(e15)), inference(backward_demodulation, [], [f1075, f1246])).
fof(f1075, plain, ~ (op2(e26, e20) = op2(e26, e26)), inference(cnf_transformation, [], [f8])).
fof(f12684, plain, (~ spl42_382 | ~ spl42_384), inference(avatar_contradiction_clause, [], [f12683])).
fof(f12683, plain, ($false | (~ spl42_382 | ~ spl42_384)), inference(subsumption_resolution, [], [f12681, f1128])).
fof(f12681, plain, ((e23 = e25) | (~ spl42_382 | ~ spl42_384)), inference(backward_demodulation, [], [f3277, f3269])).
fof(f3277, plain, ((e25 = op2(e26, e21)) | ~ spl42_384), inference(avatar_component_clause, [], [f3275])).
fof(f3275, plain, (spl42_384 <=> (e25 = op2(e26, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_384])])).
fof(f12653, plain, (~ spl42_698 | ~ spl42_375), inference(avatar_split_clause, [], [f12648, f3238, f4872])).
fof(f3238, plain, (spl42_375 <=> (e23 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_375])])).
fof(f12648, plain, (~ (e23 = h7(e13)) | ~ spl42_375), inference(backward_demodulation, [], [f4797, f3240])).
fof(f3240, plain, ((e23 = op2(e26, e22)) | ~ spl42_375), inference(avatar_component_clause, [], [f3238])).
fof(f4797, plain, ~ (op2(e26, e22) = h7(e13)), inference(backward_demodulation, [], [f1076, f4789])).
fof(f1076, plain, ~ (op2(e26, e21) = op2(e26, e22)), inference(cnf_transformation, [], [f8])).
fof(f12641, plain, (~ spl42_377 | ~ spl42_688), inference(avatar_split_clause, [], [f12640, f4816, f3246])).
fof(f3246, plain, (spl42_377 <=> (e25 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_377])])).
fof(f12640, plain, (~ (e25 = op2(e26, e22)) | ~ spl42_688), inference(forward_demodulation, [], [f4779, f4817])).
fof(f4779, plain, ~ (op2(e26, e22) = h7(e15)), inference(backward_demodulation, [], [f1084, f1246])).
fof(f1084, plain, ~ (op2(e26, e22) = op2(e26, e26)), inference(cnf_transformation, [], [f8])).
fof(f12636, plain, (~ spl42_789 | ~ spl42_369), inference(avatar_split_clause, [], [f12630, f3213, f5344])).
fof(f12630, plain, (~ (e24 = h4(e15)) | ~ spl42_369), inference(backward_demodulation, [], [f4713, f3215])).
fof(f4713, plain, ~ (op2(e26, e23) = h4(e15)), inference(backward_demodulation, [], [f877, f1225])).
fof(f877, plain, ~ (op2(e23, e23) = op2(e26, e23)), inference(cnf_transformation, [], [f8])).
fof(f12635, plain, (~ spl42_376 | ~ spl42_369), inference(avatar_split_clause, [], [f12627, f3213, f3242])).
fof(f3242, plain, (spl42_376 <=> (e24 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_376])])).
fof(f12627, plain, (~ (e24 = op2(e26, e22)) | ~ spl42_369), inference(backward_demodulation, [], [f1081, f3215])).
fof(f1081, plain, ~ (op2(e26, e22) = op2(e26, e23)), inference(cnf_transformation, [], [f8])).
fof(f12634, plain, (~ spl42_418 | ~ spl42_369), inference(avatar_split_clause, [], [f12626, f3213, f3416])).
fof(f12626, plain, (~ (e24 = op2(e25, e23)) | ~ spl42_369), inference(backward_demodulation, [], [f880, f3215])).
fof(f880, plain, ~ (op2(e25, e23) = op2(e26, e23)), inference(cnf_transformation, [], [f8])).
fof(f12633, plain, (~ spl42_614 | ~ spl42_369), inference(avatar_split_clause, [], [f12625, f3213, f4228])).
fof(f4228, plain, (spl42_614 <=> (e24 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_614])])).
fof(f12625, plain, (~ (e24 = op2(e21, e23)) | ~ spl42_369), inference(backward_demodulation, [], [f870, f3215])).
fof(f870, plain, ~ (op2(e21, e23) = op2(e26, e23)), inference(cnf_transformation, [], [f8])).
fof(f12596, plain, (~ spl42_363 | ~ spl42_688), inference(avatar_split_clause, [], [f12595, f4816, f3188])).
fof(f12595, plain, (~ (e25 = op2(e26, e24)) | ~ spl42_688), inference(forward_demodulation, [], [f4781, f4817])).
fof(f4781, plain, ~ (op2(e26, e24) = h7(e15)), inference(backward_demodulation, [], [f1089, f1246])).
fof(f1089, plain, ~ (op2(e26, e24) = op2(e26, e26)), inference(cnf_transformation, [], [f8])).
fof(f12585, plain, (~ spl42_740 | ~ spl42_351), inference(avatar_split_clause, [], [f12584, f3139, f5098])).
fof(f5098, plain, (spl42_740 <=> (e20 = h6(e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_740])])).
fof(f12584, plain, (~ (e20 = h6(e15)) | ~ spl42_351), inference(forward_demodulation, [], [f4755, f3141])).
fof(f4755, plain, ~ (op2(e26, e25) = h6(e15)), inference(backward_demodulation, [], [f922, f1239])).
fof(f922, plain, ~ (op2(e25, e25) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f12580, plain, (~ spl42_372 | ~ spl42_351), inference(avatar_split_clause, [], [f12579, f3139, f3226])).
fof(f3226, plain, (spl42_372 <=> (e20 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_372])])).
fof(f12579, plain, (~ (e20 = op2(e26, e22)) | ~ spl42_351), inference(forward_demodulation, [], [f1083, f3141])).
fof(f1083, plain, ~ (op2(e26, e22) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f12573, plain, (spl42_646 | ~ spl42_351 | ~ spl42_434), inference(avatar_split_clause, [], [f12572, f3482, f3139, f4361])).
fof(f3482, plain, (spl42_434 <=> (e26 = op2(e25, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_434])])).
fof(f12572, plain, ((e21 = op2(e20, e25)) | (~ spl42_351 | ~ spl42_434)), inference(forward_demodulation, [], [f11586, f3141])).
fof(f11586, plain, ((e21 = op2(op2(e26, e25), e25)) | ~ spl42_434), inference(backward_demodulation, [], [f466, f3484])).
fof(f3484, plain, ((e26 = op2(e25, e21)) | ~ spl42_434), inference(avatar_component_clause, [], [f3482])).
fof(f12551, plain, (~ spl42_315 | ~ spl42_301), inference(avatar_split_clause, [], [f12547, f2833, f2891])).
fof(f12547, plain, (~ (e16 = op1(e10, e14)) | ~ spl42_301), inference(backward_demodulation, [], [f669, f2835])).
fof(f669, plain, ~ (op1(e10, e14) = op1(e10, e16)), inference(cnf_transformation, [], [f7])).
fof(f12550, plain, (~ spl42_322 | ~ spl42_301), inference(avatar_split_clause, [], [f12546, f2833, f2920])).
fof(f2920, plain, (spl42_322 <=> (e16 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_322])])).
fof(f12546, plain, (~ (e16 = op1(e10, e13)) | ~ spl42_301), inference(backward_demodulation, [], [f667, f2835])).
fof(f667, plain, ~ (op1(e10, e13) = op1(e10, e16)), inference(cnf_transformation, [], [f7])).
fof(f12549, plain, (~ spl42_329 | ~ spl42_301), inference(avatar_split_clause, [], [f12545, f2833, f2949])).
fof(f2949, plain, (spl42_329 <=> (e16 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_329])])).
fof(f12545, plain, (~ (e16 = op1(e10, e12)) | ~ spl42_301), inference(backward_demodulation, [], [f664, f2835])).
fof(f664, plain, ~ (op1(e10, e12) = op1(e10, e16)), inference(cnf_transformation, [], [f7])).
fof(f12540, plain, (~ spl42_284 | ~ spl42_39), inference(avatar_split_clause, [], [f12539, f1748, f2763])).
fof(f12539, plain, (~ (e13 = op1(e11, e11)) | ~ spl42_39), inference(forward_demodulation, [], [f534, f1750])).
fof(f534, plain, ~ (op1(e11, e11) = op1(e16, e11)), inference(cnf_transformation, [], [f7])).
fof(f12524, plain, (~ spl42_268 | ~ spl42_219), inference(avatar_split_clause, [], [f12523, f2494, f2697])).
fof(f2494, plain, (spl42_219 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_219])])).
fof(f12523, plain, (~ (e11 = op1(e11, e13)) | ~ spl42_219), inference(forward_demodulation, [], [f572, f2496])).
fof(f2496, plain, ((e11 = op1(e12, e13)) | ~ spl42_219), inference(avatar_component_clause, [], [f2494])).
fof(f572, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f7])).
fof(f12520, plain, (~ spl42_269 | ~ spl42_255), inference(avatar_split_clause, [], [f12519, f2643, f2701])).
fof(f2701, plain, (spl42_269 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_269])])).
fof(f12519, plain, (~ (e12 = op1(e11, e13)) | ~ spl42_255), inference(forward_demodulation, [], [f687, f2645])).
fof(f687, plain, ~ (op1(e11, e13) = op1(e11, e15)), inference(cnf_transformation, [], [f7])).
fof(f12518, plain, (~ spl42_264 | ~ spl42_68), inference(avatar_split_clause, [], [f12517, f1868, f2680])).
fof(f1868, plain, (spl42_68 <=> (e14 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_68])])).
fof(f12517, plain, (~ (e14 = op1(e11, e14)) | ~ spl42_68), inference(forward_demodulation, [], [f596, f1870])).
fof(f1870, plain, ((e14 = op1(e15, e14)) | ~ spl42_68), inference(avatar_component_clause, [], [f1868])).
fof(f596, plain, ~ (op1(e11, e14) = op1(e15, e14)), inference(cnf_transformation, [], [f7])).
fof(f12516, plain, (~ spl42_262 | ~ spl42_255), inference(avatar_split_clause, [], [f12515, f2643, f2672])).
fof(f2672, plain, (spl42_262 <=> (e12 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_262])])).
fof(f12515, plain, (~ (e12 = op1(e11, e14)) | ~ spl42_255), inference(forward_demodulation, [], [f689, f2645])).
fof(f689, plain, ~ (op1(e11, e14) = op1(e11, e15)), inference(cnf_transformation, [], [f7])).
fof(f12493, plain, (~ spl42_251 | ~ spl42_6), inference(avatar_split_clause, [], [f10775, f1611, f2626])).
fof(f10775, plain, (~ (e15 = op1(e11, e16)) | ~ spl42_6), inference(forward_demodulation, [], [f639, f1613])).
fof(f639, plain, ~ (op1(e11, e16) = op1(e16, e16)), inference(cnf_transformation, [], [f7])).
fof(f12466, plain, (~ spl42_205 | ~ spl42_58), inference(avatar_split_clause, [], [f12465, f1827, f2436])).
fof(f2436, plain, (spl42_205 <=> (e11 = op1(e12, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_205])])).
fof(f12465, plain, (~ (e11 = op1(e12, e15)) | ~ spl42_58), inference(forward_demodulation, [], [f621, f1829])).
fof(f621, plain, ~ (op1(e12, e15) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f12441, plain, (~ spl42_202 | ~ spl42_6), inference(avatar_split_clause, [], [f10730, f1611, f2423])).
fof(f10730, plain, (~ (e15 = op1(e12, e16)) | ~ spl42_6), inference(forward_demodulation, [], [f643, f1613])).
fof(f643, plain, ~ (op1(e12, e16) = op1(e16, e16)), inference(cnf_transformation, [], [f7])).
fof(f12430, plain, (~ spl42_184 | ~ spl42_177), inference(avatar_split_clause, [], [f12429, f2320, f2349])).
fof(f2320, plain, (spl42_177 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_177])])).
fof(f12429, plain, (~ (e11 = op1(e13, e11)) | ~ spl42_177), inference(forward_demodulation, [], [f719, f2322])).
fof(f2322, plain, ((e11 = op1(e13, e12)) | ~ spl42_177), inference(avatar_component_clause, [], [f2320])).
fof(f719, plain, ~ (op1(e13, e11) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f12418, plain, (~ spl42_156 | ~ spl42_58), inference(avatar_split_clause, [], [f12417, f1827, f2233])).
fof(f2233, plain, (spl42_156 <=> (e11 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_156])])).
fof(f12417, plain, (~ (e11 = op1(e13, e15)) | ~ spl42_58), inference(forward_demodulation, [], [f624, f1829])).
fof(f624, plain, ~ (op1(e13, e15) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f12390, plain, (~ spl42_153 | ~ spl42_6), inference(avatar_split_clause, [], [f10685, f1611, f2220])).
fof(f10685, plain, (~ (e15 = op1(e13, e16)) | ~ spl42_6), inference(forward_demodulation, [], [f646, f1613])).
fof(f646, plain, ~ (op1(e13, e16) = op1(e16, e16)), inference(cnf_transformation, [], [f7])).
fof(f12381, plain, (~ spl42_138 | ~ spl42_110), inference(avatar_split_clause, [], [f12380, f2042, f2158])).
fof(f2158, plain, (spl42_138 <=> (e14 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_138])])).
fof(f2042, plain, (spl42_110 <=> (e14 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_110])])).
fof(f12380, plain, (~ (e14 = op1(e14, e11)) | ~ spl42_110), inference(forward_demodulation, [], [f743, f2044])).
fof(f2044, plain, ((e14 = op1(e14, e15)) | ~ spl42_110), inference(avatar_component_clause, [], [f2042])).
fof(f743, plain, ~ (op1(e14, e11) = op1(e14, e15)), inference(cnf_transformation, [], [f7])).
fof(f12361, plain, (~ spl42_107 | ~ spl42_58), inference(avatar_split_clause, [], [f12360, f1827, f2030])).
fof(f2030, plain, (spl42_107 <=> (e11 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_107])])).
fof(f12360, plain, (~ (e11 = op1(e14, e15)) | ~ spl42_58), inference(forward_demodulation, [], [f626, f1829])).
fof(f626, plain, ~ (op1(e14, e15) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f12336, plain, (~ spl42_104 | ~ spl42_6), inference(avatar_split_clause, [], [f10646, f1611, f2017])).
fof(f10646, plain, (~ (e15 = op1(e14, e16)) | ~ spl42_6), inference(forward_demodulation, [], [f648, f1613])).
fof(f648, plain, ~ (op1(e14, e16) = op1(e16, e16)), inference(cnf_transformation, [], [f7])).
fof(f12322, plain, (spl42_209 | ~ spl42_58 | ~ spl42_255), inference(avatar_split_clause, [], [f12321, f2643, f1827, f2452])).
fof(f12321, plain, ((e15 = op1(e12, e15)) | (~ spl42_58 | ~ spl42_255)), inference(forward_demodulation, [], [f12320, f2645])).
fof(f12320, plain, ((e15 = op1(op1(e11, e15), e15)) | ~ spl42_58), inference(forward_demodulation, [], [f298, f1829])).
fof(f298, plain, (e15 = op1(op1(op1(e15, e15), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f12303, plain, (~ spl42_58 | ~ spl42_76 | ~ spl42_255), inference(avatar_contradiction_clause, [], [f12302])).
fof(f12302, plain, ($false | (~ spl42_58 | ~ spl42_76 | ~ spl42_255)), inference(subsumption_resolution, [], [f12301, f1102])).
fof(f12301, plain, ((e12 = e13) | (~ spl42_58 | ~ spl42_76 | ~ spl42_255)), inference(forward_demodulation, [], [f12300, f2645])).
fof(f12300, plain, ((e13 = op1(e11, e15)) | (~ spl42_58 | ~ spl42_76)), inference(forward_demodulation, [], [f12299, f1829])).
fof(f12299, plain, ((e13 = op1(op1(e15, e15), e15)) | ~ spl42_76), inference(forward_demodulation, [], [f284, f1903])).
fof(f1903, plain, ((e15 = op1(e15, e13)) | ~ spl42_76), inference(avatar_component_clause, [], [f1901])).
fof(f1901, plain, (spl42_76 <=> (e15 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_76])])).
fof(f12296, plain, (~ spl42_72 | ~ spl42_58), inference(avatar_split_clause, [], [f12295, f1827, f1885])).
fof(f1885, plain, (spl42_72 <=> (e11 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_72])])).
fof(f12295, plain, (~ (e11 = op1(e15, e13)) | ~ spl42_58), inference(forward_demodulation, [], [f771, f1829])).
fof(f771, plain, ~ (op1(e15, e13) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f12293, plain, (~ spl42_65 | ~ spl42_58), inference(avatar_split_clause, [], [f12292, f1827, f1856])).
fof(f1856, plain, (spl42_65 <=> (e11 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_65])])).
fof(f12292, plain, (~ (e11 = op1(e15, e14)) | ~ spl42_58), inference(forward_demodulation, [], [f773, f1829])).
fof(f773, plain, ~ (op1(e15, e14) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f12266, plain, (~ spl42_55 | ~ spl42_6), inference(avatar_split_clause, [], [f10612, f1611, f1814])).
fof(f1814, plain, (spl42_55 <=> (e15 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_55])])).
fof(f10612, plain, (~ (e15 = op1(e15, e16)) | ~ spl42_6), inference(forward_demodulation, [], [f649, f1613])).
fof(f649, plain, ~ (op1(e15, e16) = op1(e16, e16)), inference(cnf_transformation, [], [f7])).
fof(f12265, plain, (~ spl42_51 | ~ spl42_58), inference(avatar_split_clause, [], [f12264, f1827, f1798])).
fof(f1798, plain, (spl42_51 <=> (e11 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_51])])).
fof(f12264, plain, (~ (e11 = op1(e15, e16)) | ~ spl42_58), inference(forward_demodulation, [], [f775, f1829])).
fof(f775, plain, ~ (op1(e15, e15) = op1(e15, e16)), inference(cnf_transformation, [], [f7])).
fof(f12254, plain, (~ spl42_48 | ~ spl42_6), inference(avatar_split_clause, [], [f12253, f1611, f1785])).
fof(f1785, plain, (spl42_48 <=> (e15 = op1(e16, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_48])])).
fof(f12253, plain, (~ (e15 = op1(e16, e10)) | ~ spl42_6), inference(forward_demodulation, [], [f781, f1613])).
fof(f781, plain, ~ (op1(e16, e10) = op1(e16, e16)), inference(cnf_transformation, [], [f7])).
fof(f12244, plain, (~ spl42_88 | ~ spl42_39), inference(avatar_split_clause, [], [f12237, f1748, f1951])).
fof(f12237, plain, (~ (e13 = op1(e15, e11)) | ~ spl42_39), inference(backward_demodulation, [], [f544, f1750])).
fof(f544, plain, ~ (op1(e15, e11) = op1(e16, e11)), inference(cnf_transformation, [], [f7])).
fof(f12243, plain, (~ spl42_137 | ~ spl42_39), inference(avatar_split_clause, [], [f12236, f1748, f2154])).
fof(f12236, plain, (~ (e13 = op1(e14, e11)) | ~ spl42_39), inference(backward_demodulation, [], [f543, f1750])).
fof(f543, plain, ~ (op1(e14, e11) = op1(e16, e11)), inference(cnf_transformation, [], [f7])).
fof(f12242, plain, (~ spl42_186 | ~ spl42_39), inference(avatar_split_clause, [], [f12235, f1748, f2357])).
fof(f12235, plain, (~ (e13 = op1(e13, e11)) | ~ spl42_39), inference(backward_demodulation, [], [f541, f1750])).
fof(f541, plain, ~ (op1(e13, e11) = op1(e16, e11)), inference(cnf_transformation, [], [f7])).
fof(f12241, plain, (~ spl42_235 | ~ spl42_39), inference(avatar_split_clause, [], [f12234, f1748, f2560])).
fof(f12234, plain, (~ (e13 = op1(e12, e11)) | ~ spl42_39), inference(backward_demodulation, [], [f538, f1750])).
fof(f538, plain, ~ (op1(e12, e11) = op1(e16, e11)), inference(cnf_transformation, [], [f7])).
fof(f12223, plain, (~ spl42_33 | ~ spl42_26), inference(avatar_split_clause, [], [f12215, f1694, f1723])).
fof(f12215, plain, (~ (e14 = op1(e16, e12)) | ~ spl42_26), inference(backward_demodulation, [], [f787, f1696])).
fof(f787, plain, ~ (op1(e16, e12) = op1(e16, e13)), inference(cnf_transformation, [], [f7])).
fof(f12221, plain, (~ spl42_75 | ~ spl42_26), inference(avatar_split_clause, [], [f12213, f1694, f1897])).
fof(f12213, plain, (~ (e14 = op1(e15, e13)) | ~ spl42_26), inference(backward_demodulation, [], [f586, f1696])).
fof(f586, plain, ~ (op1(e15, e13) = op1(e16, e13)), inference(cnf_transformation, [], [f7])).
fof(f12220, plain, (~ spl42_124 | ~ spl42_26), inference(avatar_split_clause, [], [f12212, f1694, f2100])).
fof(f12212, plain, (~ (e14 = op1(e14, e13)) | ~ spl42_26), inference(backward_demodulation, [], [f585, f1696])).
fof(f585, plain, ~ (op1(e14, e13) = op1(e16, e13)), inference(cnf_transformation, [], [f7])).
fof(f12219, plain, (~ spl42_271 | ~ spl42_26), inference(avatar_split_clause, [], [f12211, f1694, f2709])).
fof(f2709, plain, (spl42_271 <=> (e14 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_271])])).
fof(f12211, plain, (~ (e14 = op1(e11, e13)) | ~ spl42_26), inference(backward_demodulation, [], [f576, f1696])).
fof(f576, plain, ~ (op1(e11, e13) = op1(e16, e13)), inference(cnf_transformation, [], [f7])).
fof(f12218, plain, (~ spl42_320 | ~ spl42_26), inference(avatar_split_clause, [], [f12210, f1694, f2912])).
fof(f2912, plain, (spl42_320 <=> (e14 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_320])])).
fof(f12210, plain, (~ (e14 = op1(e10, e13)) | ~ spl42_26), inference(backward_demodulation, [], [f571, f1696])).
fof(f571, plain, ~ (op1(e10, e13) = op1(e16, e13)), inference(cnf_transformation, [], [f7])).
fof(f12207, plain, (spl42_8 | ~ spl42_6), inference(avatar_split_clause, [], [f10511, f1611, f1620])).
fof(f10511, plain, ((e10 = op1(e16, e15)) | ~ spl42_6), inference(backward_demodulation, [], [f1186, f1613])).
fof(f1186, plain, (e10 = op1(e16, op1(e16, e16))), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((e15 = op1(e16, e16)) & (e14 = op1(e16, op1(e16, op1(op1(e16, e16), op1(e16, e16))))) & (e13 = op1(e16, op1(op1(e16, e16), op1(e16, e16)))) & (e12 = op1(op1(op1(e16, e16), op1(e16, e16)), op1(e16, e16))) & (e11 = op1(op1(e16, e16), op1(e16, e16))) & (e10 = op1(e16, op1(e16, e16)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax14)).
fof(f12183, plain, (~ spl42_10 | ~ spl42_255), inference(avatar_split_clause, [], [f12182, f2643, f1628])).
fof(f1628, plain, (spl42_10 <=> (e12 = op1(e16, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_10])])).
fof(f12182, plain, (~ (e12 = op1(e16, e15)) | ~ spl42_255), inference(forward_demodulation, [], [f618, f2645])).
fof(f618, plain, ~ (op1(e11, e15) = op1(e16, e15)), inference(cnf_transformation, [], [f7])).
fof(f12159, plain, (~ spl42_58 | ~ spl42_97 | spl42_253), inference(avatar_contradiction_clause, [], [f12158])).
fof(f12158, plain, ($false | (~ spl42_58 | ~ spl42_97 | spl42_253)), inference(subsumption_resolution, [], [f12157, f2636])).
fof(f2636, plain, (~ (e10 = op1(e11, e15)) | spl42_253), inference(avatar_component_clause, [], [f2635])).
fof(f2635, plain, (spl42_253 <=> (e10 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_253])])).
fof(f12157, plain, ((e10 = op1(e11, e15)) | (~ spl42_58 | ~ spl42_97)), inference(backward_demodulation, [], [f12081, f1829])).
fof(f12081, plain, ((e10 = op1(op1(e15, e15), e15)) | ~ spl42_97), inference(backward_demodulation, [], [f263, f1990])).
fof(f1990, plain, ((e15 = op1(e15, e10)) | ~ spl42_97), inference(avatar_component_clause, [], [f1988])).
fof(f1988, plain, (spl42_97 <=> (e15 = op1(e15, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_97])])).
fof(f12153, plain, (~ spl42_68 | ~ spl42_70), inference(avatar_contradiction_clause, [], [f12152])).
fof(f12152, plain, ($false | (~ spl42_68 | ~ spl42_70)), inference(subsumption_resolution, [], [f12151, f1110])).
fof(f12151, plain, ((e14 = e16) | (~ spl42_68 | ~ spl42_70)), inference(backward_demodulation, [], [f1878, f1870])).
fof(f12122, plain, (spl42_304 | ~ spl42_8 | ~ spl42_84), inference(avatar_split_clause, [], [f12121, f1934, f1620, f2846])).
fof(f1934, plain, (spl42_84 <=> (e16 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_84])])).
fof(f12121, plain, ((e12 = op1(e10, e15)) | (~ spl42_8 | ~ spl42_84)), inference(forward_demodulation, [], [f12117, f1622])).
fof(f12117, plain, ((e12 = op1(op1(e16, e15), e15)) | ~ spl42_84), inference(backward_demodulation, [], [f277, f1936])).
fof(f1936, plain, ((e16 = op1(e15, e12)) | ~ spl42_84), inference(avatar_component_clause, [], [f1934])).
fof(f277, plain, (e12 = op1(op1(op1(e15, e12), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f12116, plain, (~ spl42_87 | ~ spl42_89), inference(avatar_contradiction_clause, [], [f12115])).
fof(f12115, plain, ($false | (~ spl42_87 | ~ spl42_89)), inference(subsumption_resolution, [], [f12114, f1103])).
fof(f12114, plain, ((e12 = e14) | (~ spl42_87 | ~ spl42_89)), inference(backward_demodulation, [], [f1957, f1949])).
fof(f1957, plain, ((e14 = op1(e15, e11)) | ~ spl42_89), inference(avatar_component_clause, [], [f1955])).
fof(f1955, plain, (spl42_89 <=> (e14 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_89])])).
fof(f12106, plain, (spl42_303 | ~ spl42_8 | ~ spl42_91), inference(avatar_split_clause, [], [f12105, f1963, f1620, f2842])).
fof(f1963, plain, (spl42_91 <=> (e16 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_91])])).
fof(f12105, plain, ((e11 = op1(e10, e15)) | (~ spl42_8 | ~ spl42_91)), inference(forward_demodulation, [], [f12100, f1622])).
fof(f12100, plain, ((e11 = op1(op1(e16, e15), e15)) | ~ spl42_91), inference(backward_demodulation, [], [f270, f1965])).
fof(f1965, plain, ((e16 = op1(e15, e11)) | ~ spl42_91), inference(avatar_component_clause, [], [f1963])).
fof(f12000, plain, (~ spl42_163 | ~ spl42_268), inference(avatar_contradiction_clause, [], [f11999])).
fof(f11999, plain, ($false | (~ spl42_163 | ~ spl42_268)), inference(subsumption_resolution, [], [f11998, f1099])).
fof(f11998, plain, ((e11 = e14) | (~ spl42_163 | ~ spl42_268)), inference(forward_demodulation, [], [f11997, f2699])).
fof(f11997, plain, ((e14 = op1(e11, e13)) | (~ spl42_163 | ~ spl42_268)), inference(forward_demodulation, [], [f11994, f2699])).
fof(f11994, plain, ((e14 = op1(op1(e11, e13), e13)) | ~ spl42_163), inference(backward_demodulation, [], [f289, f2264])).
fof(f2264, plain, ((e11 = op1(e13, e14)) | ~ spl42_163), inference(avatar_component_clause, [], [f2262])).
fof(f2262, plain, (spl42_163 <=> (e11 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_163])])).
fof(f11989, plain, (~ spl42_177 | ~ spl42_268), inference(avatar_contradiction_clause, [], [f11988])).
fof(f11988, plain, ($false | (~ spl42_177 | ~ spl42_268)), inference(subsumption_resolution, [], [f11987, f1097])).
fof(f11987, plain, ((e11 = e12) | (~ spl42_177 | ~ spl42_268)), inference(forward_demodulation, [], [f11986, f2699])).
fof(f11986, plain, ((e12 = op1(e11, e13)) | (~ spl42_177 | ~ spl42_268)), inference(forward_demodulation, [], [f11981, f2699])).
fof(f11981, plain, ((e12 = op1(op1(e11, e13), e13)) | ~ spl42_177), inference(backward_demodulation, [], [f275, f2322])).
fof(f11965, plain, (~ spl42_184 | ~ spl42_191), inference(avatar_split_clause, [], [f11956, f2378, f2349])).
fof(f11956, plain, (~ (e11 = op1(e13, e11)) | ~ spl42_191), inference(backward_demodulation, [], [f713, f2380])).
fof(f2380, plain, ((e11 = op1(e13, e10)) | ~ spl42_191), inference(avatar_component_clause, [], [f2378])).
fof(f713, plain, ~ (op1(e13, e10) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f11962, plain, (~ spl42_191 | spl42_267 | ~ spl42_268), inference(avatar_contradiction_clause, [], [f11961])).
fof(f11961, plain, ($false | (~ spl42_191 | spl42_267 | ~ spl42_268)), inference(subsumption_resolution, [], [f11960, f2694])).
fof(f2694, plain, (~ (e10 = op1(e11, e13)) | spl42_267), inference(avatar_component_clause, [], [f2693])).
fof(f11960, plain, ((e10 = op1(e11, e13)) | (~ spl42_191 | ~ spl42_268)), inference(forward_demodulation, [], [f11953, f2699])).
fof(f11953, plain, ((e10 = op1(op1(e11, e13), e13)) | ~ spl42_191), inference(backward_demodulation, [], [f261, f2380])).
fof(f261, plain, (e10 = op1(op1(op1(e13, e10), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f11903, plain, (~ spl42_184 | ~ spl42_233), inference(avatar_split_clause, [], [f11896, f2552, f2349])).
fof(f11896, plain, (~ (e11 = op1(e13, e11)) | ~ spl42_233), inference(backward_demodulation, [], [f535, f2554])).
fof(f2554, plain, ((e11 = op1(e12, e11)) | ~ spl42_233), inference(avatar_component_clause, [], [f2552])).
fof(f535, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f11854, plain, (~ spl42_100 | ~ spl42_261), inference(avatar_contradiction_clause, [], [f11853])).
fof(f11853, plain, ($false | (~ spl42_100 | ~ spl42_261)), inference(subsumption_resolution, [], [f11852, f1101])).
fof(f11852, plain, ((e11 = e16) | (~ spl42_100 | ~ spl42_261)), inference(forward_demodulation, [], [f11843, f2670])).
fof(f2670, plain, ((e11 = op1(e11, e14)) | ~ spl42_261), inference(avatar_component_clause, [], [f2668])).
fof(f11843, plain, ((e16 = op1(e11, e14)) | (~ spl42_100 | ~ spl42_261)), inference(backward_demodulation, [], [f10650, f2670])).
fof(f10650, plain, ((e16 = op1(op1(e11, e14), e14)) | ~ spl42_100), inference(forward_demodulation, [], [f304, f2003])).
fof(f11812, plain, (~ spl42_268 | ~ spl42_275), inference(avatar_split_clause, [], [f11801, f2726, f2697])).
fof(f2726, plain, (spl42_275 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_275])])).
fof(f11801, plain, (~ (e11 = op1(e11, e13)) | ~ spl42_275), inference(backward_demodulation, [], [f682, f2728])).
fof(f2728, plain, ((e11 = op1(e11, e12)) | ~ spl42_275), inference(avatar_component_clause, [], [f2726])).
fof(f682, plain, ~ (op1(e11, e12) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f11735, plain, (~ spl42_268 | ~ spl42_317), inference(avatar_split_clause, [], [f11725, f2900, f2697])).
fof(f2900, plain, (spl42_317 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_317])])).
fof(f11725, plain, (~ (e11 = op1(e11, e13)) | ~ spl42_317), inference(backward_demodulation, [], [f566, f2902])).
fof(f2902, plain, ((e11 = op1(e10, e13)) | ~ spl42_317), inference(avatar_component_clause, [], [f2900])).
fof(f566, plain, ~ (op1(e10, e13) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f11694, plain, (~ spl42_184 | ~ spl42_331), inference(avatar_split_clause, [], [f11685, f2958, f2349])).
fof(f11685, plain, (~ (e11 = op1(e13, e11)) | ~ spl42_331), inference(backward_demodulation, [], [f526, f2960])).
fof(f11682, plain, (~ spl42_351 | ~ spl42_352), inference(avatar_contradiction_clause, [], [f11681])).
fof(f11681, plain, ($false | (~ spl42_351 | ~ spl42_352)), inference(subsumption_resolution, [], [f11680, f1112])).
fof(f11680, plain, ((e20 = e21) | (~ spl42_351 | ~ spl42_352)), inference(backward_demodulation, [], [f3145, f3141])).
fof(f3145, plain, ((e21 = op2(e26, e25)) | ~ spl42_352), inference(avatar_component_clause, [], [f3143])).
fof(f3143, plain, (spl42_352 <=> (e21 = op2(e26, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_352])])).
fof(f11656, plain, (spl42_733 | ~ spl42_401), inference(avatar_split_clause, [], [f11655, f3346, f5063])).
fof(f11655, plain, ((e21 = h6(e15)) | ~ spl42_401), inference(backward_demodulation, [], [f1239, f3348])).
fof(f11650, plain, (~ spl42_411 | ~ spl42_413), inference(avatar_contradiction_clause, [], [f11649])).
fof(f11649, plain, ($false | (~ spl42_411 | ~ spl42_413)), inference(subsumption_resolution, [], [f11648, f1131])).
fof(f11648, plain, ((e24 = e26) | (~ spl42_411 | ~ spl42_413)), inference(backward_demodulation, [], [f3397, f3389])).
fof(f11605, plain, (~ spl42_430 | ~ spl42_432), inference(avatar_contradiction_clause, [], [f11604])).
fof(f11604, plain, ($false | (~ spl42_430 | ~ spl42_432)), inference(subsumption_resolution, [], [f11603, f1124])).
fof(f11603, plain, ((e22 = e24) | (~ spl42_430 | ~ spl42_432)), inference(backward_demodulation, [], [f3476, f3468])).
fof(f3476, plain, ((e24 = op2(e25, e21)) | ~ spl42_432), inference(avatar_component_clause, [], [f3474])).
fof(f3474, plain, (spl42_432 <=> (e24 = op2(e25, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_432])])).
fof(f11500, plain, (~ spl42_502 | ~ spl42_504), inference(avatar_contradiction_clause, [], [f11499])).
fof(f11499, plain, ($false | (~ spl42_502 | ~ spl42_504)), inference(subsumption_resolution, [], [f11498, f1131])).
fof(f11498, plain, ((e24 = e26) | (~ spl42_502 | ~ spl42_504)), inference(backward_demodulation, [], [f3774, f3766])).
fof(f11486, plain, (spl42_803 | ~ spl42_513), inference(avatar_split_clause, [], [f11485, f3810, f5415])).
fof(f3810, plain, (spl42_513 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_513])])).
fof(f11485, plain, ((e21 = h4(e15)) | ~ spl42_513), inference(backward_demodulation, [], [f1225, f3812])).
fof(f3812, plain, ((e21 = op2(e23, e23)) | ~ spl42_513), inference(avatar_component_clause, [], [f3810])).
fof(f11365, plain, (~ spl42_598 | ~ spl42_600), inference(avatar_contradiction_clause, [], [f11364])).
fof(f11364, plain, ($false | (~ spl42_598 | ~ spl42_600)), inference(subsumption_resolution, [], [f11363, f1124])).
fof(f11363, plain, ((e22 = e24) | (~ spl42_598 | ~ spl42_600)), inference(backward_demodulation, [], [f4172, f4164])).
fof(f4172, plain, ((e24 = op2(e21, e25)) | ~ spl42_600), inference(avatar_component_clause, [], [f4170])).
fof(f4170, plain, (spl42_600 <=> (e24 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_600])])).
fof(f11340, plain, (~ spl42_443 | ~ spl42_604), inference(avatar_contradiction_clause, [], [f11339])).
fof(f11339, plain, ($false | (~ spl42_443 | ~ spl42_604)), inference(subsumption_resolution, [], [f11338, f1122])).
fof(f11338, plain, ((e21 = e26) | (~ spl42_443 | ~ spl42_604)), inference(forward_demodulation, [], [f11329, f4189])).
fof(f4189, plain, ((e21 = op2(e21, e24)) | ~ spl42_604), inference(avatar_component_clause, [], [f4187])).
fof(f4187, plain, (spl42_604 <=> (e21 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_604])])).
fof(f11329, plain, ((e26 = op2(e21, e24)) | (~ spl42_443 | ~ spl42_604)), inference(backward_demodulation, [], [f10996, f4189])).
fof(f10996, plain, ((e26 = op2(op2(e21, e24), e24)) | ~ spl42_443), inference(forward_demodulation, [], [f500, f3522])).
fof(f11207, plain, (spl42_626 | ~ spl42_617 | ~ spl42_674), inference(avatar_split_clause, [], [f11198, f4477, f4241, f4278])).
fof(f4241, plain, (spl42_617 <=> (e20 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_617])])).
fof(f11198, plain, ((e22 = op2(e21, e21)) | (~ spl42_617 | ~ spl42_674)), inference(backward_demodulation, [], [f8764, f4479])).
fof(f8764, plain, ((e22 = op2(op2(e20, e21), e21)) | ~ spl42_617), inference(backward_demodulation, [], [f469, f4243])).
fof(f4243, plain, ((e20 = op2(e21, e22)) | ~ spl42_617), inference(avatar_component_clause, [], [f4241])).
fof(f11187, plain, (spl42_393 | ~ spl42_392 | ~ spl42_688), inference(avatar_split_clause, [], [f11182, f4816, f3308, f3313])).
fof(f11182, plain, ((e20 = op2(e25, e26)) | (~ spl42_392 | ~ spl42_688)), inference(backward_demodulation, [], [f9358, f4817])).
fof(f9358, plain, ((e20 = op2(h7(e15), e26)) | ~ spl42_392), inference(forward_demodulation, [], [f9350, f1246])).
fof(f9350, plain, ((e20 = op2(op2(e26, e26), e26)) | ~ spl42_392), inference(backward_demodulation, [], [f460, f3310])).
fof(f11185, plain, (spl42_598 | ~ spl42_688), inference(avatar_split_clause, [], [f11179, f4816, f4162])).
fof(f11179, plain, ((e22 = op2(e21, e25)) | ~ spl42_688), inference(backward_demodulation, [], [f4785, f4817])).
fof(f4785, plain, (e22 = op2(e21, h7(e15))), inference(backward_demodulation, [], [f4641, f1246])).
fof(f4641, plain, (e22 = op2(e21, op2(e26, e26))), inference(backward_demodulation, [], [f1194, f1193])).
fof(f1194, plain, (e22 = op2(op2(op2(e26, e26), op2(e26, e26)), op2(e26, e26))), inference(cnf_transformation, [], [f15])).
fof(f11184, plain, (spl42_401 | ~ spl42_688), inference(avatar_split_clause, [], [f11178, f4816, f3346])).
fof(f11178, plain, ((e21 = op2(e25, e25)) | ~ spl42_688), inference(backward_demodulation, [], [f4784, f4817])).
fof(f11183, plain, (spl42_351 | ~ spl42_688), inference(avatar_split_clause, [], [f11177, f4816, f3139])).
fof(f11177, plain, ((e20 = op2(e26, e25)) | ~ spl42_688), inference(backward_demodulation, [], [f4783, f4817])).
fof(f4783, plain, (e20 = op2(e26, h7(e15))), inference(backward_demodulation, [], [f1192, f1246])).
fof(f1192, plain, (e20 = op2(e26, op2(e26, e26))), inference(cnf_transformation, [], [f15])).
fof(f11150, plain, (~ spl42_640 | ~ spl42_644), inference(avatar_contradiction_clause, [], [f11149])).
fof(f11149, plain, ($false | (~ spl42_640 | ~ spl42_644)), inference(subsumption_resolution, [], [f11148, f1126])).
fof(f11148, plain, ((e22 = e26) | (~ spl42_640 | ~ spl42_644)), inference(forward_demodulation, [], [f4354, f4338])).
fof(f4338, plain, ((e22 = op2(e20, e26)) | ~ spl42_640), inference(avatar_component_clause, [], [f4336])).
fof(f11144, plain, (spl42_644 | ~ spl42_392), inference(avatar_split_clause, [], [f9359, f3308, f4352])).
fof(f9359, plain, ((e26 = op2(e20, e26)) | ~ spl42_392), inference(backward_demodulation, [], [f4770, f9358])).
fof(f11115, plain, (~ spl42_492 | ~ spl42_590), inference(avatar_contradiction_clause, [], [f11114])).
fof(f11114, plain, ($false | (~ spl42_492 | ~ spl42_590)), inference(subsumption_resolution, [], [f11113, f4131])).
fof(f4131, plain, ((e21 = op2(e21, e26)) | ~ spl42_590), inference(avatar_component_clause, [], [f4129])).
fof(f11113, plain, (~ (e21 = op2(e21, e26)) | ~ spl42_492), inference(forward_demodulation, [], [f930, f3725])).
fof(f930, plain, ~ (op2(e21, e26) = op2(e23, e26)), inference(cnf_transformation, [], [f8])).
fof(f11012, plain, (~ spl42_485 | ~ spl42_443), inference(avatar_split_clause, [], [f11011, f3520, f3694])).
fof(f11011, plain, (~ (e21 = op2(e24, e20)) | ~ spl42_443), inference(forward_demodulation, [], [f1033, f3522])).
fof(f1033, plain, ~ (op2(e24, e20) = op2(e24, e26)), inference(cnf_transformation, [], [f8])).
fof(f10963, plain, (~ spl42_733 | ~ spl42_394), inference(avatar_split_clause, [], [f10962, f3317, f5063])).
fof(f3317, plain, (spl42_394 <=> (e21 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_394])])).
fof(f10962, plain, (~ (e21 = h6(e15)) | ~ spl42_394), inference(forward_demodulation, [], [f4761, f3319])).
fof(f3319, plain, ((e21 = op2(e25, e26)) | ~ spl42_394), inference(avatar_component_clause, [], [f3317])).
fof(f4761, plain, ~ (op2(e25, e26) = h6(e15)), inference(backward_demodulation, [], [f1069, f1239])).
fof(f1069, plain, ~ (op2(e25, e25) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f10941, plain, (~ spl42_380 | ~ spl42_382), inference(avatar_contradiction_clause, [], [f10940])).
fof(f10940, plain, ($false | (~ spl42_380 | ~ spl42_382)), inference(subsumption_resolution, [], [f10939, f1119])).
fof(f10939, plain, ((e21 = e23) | (~ spl42_380 | ~ spl42_382)), inference(forward_demodulation, [], [f3269, f3261])).
fof(f3261, plain, ((e21 = op2(e26, e21)) | ~ spl42_380), inference(avatar_component_clause, [], [f3259])).
fof(f3259, plain, (spl42_380 <=> (e21 = op2(e26, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_380])])).
fof(f10938, plain, (~ spl42_373 | ~ spl42_590), inference(avatar_contradiction_clause, [], [f10937])).
fof(f10937, plain, ($false | (~ spl42_373 | ~ spl42_590)), inference(subsumption_resolution, [], [f10936, f1118])).
fof(f10936, plain, ((e21 = e22) | (~ spl42_373 | ~ spl42_590)), inference(forward_demodulation, [], [f10935, f4131])).
fof(f10935, plain, ((e22 = op2(e21, e26)) | (~ spl42_373 | ~ spl42_590)), inference(forward_demodulation, [], [f10934, f4131])).
fof(f10934, plain, ((e22 = op2(op2(e21, e26), e26)) | ~ spl42_373), inference(forward_demodulation, [], [f474, f3232])).
fof(f10898, plain, (~ spl42_367 | ~ spl42_369), inference(avatar_contradiction_clause, [], [f10897])).
fof(f10897, plain, ($false | (~ spl42_367 | ~ spl42_369)), inference(subsumption_resolution, [], [f10896, f1124])).
fof(f10896, plain, ((e22 = e24) | (~ spl42_367 | ~ spl42_369)), inference(forward_demodulation, [], [f3215, f3207])).
fof(f3207, plain, ((e22 = op2(e26, e23)) | ~ spl42_367), inference(avatar_component_clause, [], [f3205])).
fof(f3205, plain, (spl42_367 <=> (e22 = op2(e26, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_367])])).
fof(f10895, plain, (~ spl42_359 | ~ spl42_590), inference(avatar_contradiction_clause, [], [f10894])).
fof(f10894, plain, ($false | (~ spl42_359 | ~ spl42_590)), inference(subsumption_resolution, [], [f10893, f1120])).
fof(f10893, plain, ((e21 = e24) | (~ spl42_359 | ~ spl42_590)), inference(forward_demodulation, [], [f10892, f4131])).
fof(f10892, plain, ((e24 = op2(e21, e26)) | (~ spl42_359 | ~ spl42_590)), inference(forward_demodulation, [], [f10891, f4131])).
fof(f10891, plain, ((e24 = op2(op2(e21, e26), e26)) | ~ spl42_359), inference(forward_demodulation, [], [f488, f3174])).
fof(f10858, plain, (~ spl42_330 | ~ spl42_281), inference(avatar_split_clause, [], [f10857, f2751, f2954])).
fof(f10857, plain, (~ (e10 = op1(e10, e11)) | ~ spl42_281), inference(forward_demodulation, [], [f524, f2753])).
fof(f524, plain, ~ (op1(e10, e11) = op1(e11, e11)), inference(cnf_transformation, [], [f7])).
fof(f10791, plain, (spl42_255 | ~ spl42_6), inference(avatar_split_clause, [], [f10514, f1611, f2643])).
fof(f10514, plain, ((e12 = op1(e11, e15)) | ~ spl42_6), inference(backward_demodulation, [], [f4635, f1613])).
fof(f4635, plain, (e12 = op1(e11, op1(e16, e16))), inference(backward_demodulation, [], [f1188, f1187])).
fof(f1187, plain, (e11 = op1(op1(e16, e16), op1(e16, e16))), inference(cnf_transformation, [], [f14])).
fof(f1188, plain, (e12 = op1(op1(op1(e16, e16), op1(e16, e16)), op1(e16, e16))), inference(cnf_transformation, [], [f14])).
fof(f10787, plain, (~ spl42_198 | ~ spl42_247), inference(avatar_contradiction_clause, [], [f10786])).
fof(f10786, plain, ($false | (~ spl42_198 | ~ spl42_247)), inference(subsumption_resolution, [], [f10785, f2612])).
fof(f2612, plain, ((e11 = op1(e11, e16)) | ~ spl42_247), inference(avatar_component_clause, [], [f2610])).
fof(f10785, plain, (~ (e11 = op1(e11, e16)) | ~ spl42_198), inference(forward_demodulation, [], [f635, f2409])).
fof(f2409, plain, ((e11 = op1(e12, e16)) | ~ spl42_198), inference(avatar_component_clause, [], [f2407])).
fof(f635, plain, ~ (op1(e11, e16) = op1(e12, e16)), inference(cnf_transformation, [], [f7])).
fof(f10784, plain, (~ spl42_149 | ~ spl42_247), inference(avatar_contradiction_clause, [], [f10783])).
fof(f10783, plain, ($false | (~ spl42_149 | ~ spl42_247)), inference(subsumption_resolution, [], [f10782, f2612])).
fof(f10782, plain, (~ (e11 = op1(e11, e16)) | ~ spl42_149), inference(forward_demodulation, [], [f636, f2206])).
fof(f636, plain, ~ (op1(e11, e16) = op1(e13, e16)), inference(cnf_transformation, [], [f7])).
fof(f10778, plain, (~ spl42_246 | ~ spl42_50), inference(avatar_split_clause, [], [f10777, f1794, f2606])).
fof(f2606, plain, (spl42_246 <=> (e10 = op1(e11, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_246])])).
fof(f10777, plain, (~ (e10 = op1(e11, e16)) | ~ spl42_50), inference(forward_demodulation, [], [f638, f1796])).
fof(f638, plain, ~ (op1(e11, e16) = op1(e15, e16)), inference(cnf_transformation, [], [f7])).
fof(f10771, plain, (~ spl42_247 | ~ spl42_249), inference(avatar_contradiction_clause, [], [f10770])).
fof(f10770, plain, ($false | (~ spl42_247 | ~ spl42_249)), inference(subsumption_resolution, [], [f10769, f1098])).
fof(f10769, plain, ((e11 = e13) | (~ spl42_247 | ~ spl42_249)), inference(forward_demodulation, [], [f2620, f2612])).
fof(f10748, plain, (~ spl42_212 | ~ spl42_16), inference(avatar_split_clause, [], [f10747, f1653, f2465])).
fof(f2465, plain, (spl42_212 <=> (e11 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_212])])).
fof(f10747, plain, (~ (e11 = op1(e12, e14)) | ~ spl42_16), inference(forward_demodulation, [], [f601, f1655])).
fof(f601, plain, ~ (op1(e12, e14) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f10736, plain, (~ spl42_100 | ~ spl42_198), inference(avatar_contradiction_clause, [], [f10735])).
fof(f10735, plain, ($false | (~ spl42_100 | ~ spl42_198)), inference(subsumption_resolution, [], [f10734, f2409])).
fof(f10734, plain, (~ (e11 = op1(e12, e16)) | ~ spl42_100), inference(forward_demodulation, [], [f641, f2003])).
fof(f641, plain, ~ (op1(e12, e16) = op1(e14, e16)), inference(cnf_transformation, [], [f7])).
fof(f10733, plain, (~ spl42_197 | ~ spl42_50), inference(avatar_split_clause, [], [f10732, f1794, f2403])).
fof(f10732, plain, (~ (e10 = op1(e12, e16)) | ~ spl42_50), inference(forward_demodulation, [], [f642, f1796])).
fof(f642, plain, ~ (op1(e12, e16) = op1(e15, e16)), inference(cnf_transformation, [], [f7])).
fof(f10700, plain, (~ spl42_163 | ~ spl42_16), inference(avatar_split_clause, [], [f10699, f1653, f2262])).
fof(f10699, plain, (~ (e11 = op1(e13, e14)) | ~ spl42_16), inference(forward_demodulation, [], [f604, f1655])).
fof(f604, plain, ~ (op1(e13, e14) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f10696, plain, (~ spl42_155 | ~ spl42_8), inference(avatar_split_clause, [], [f10695, f1620, f2229])).
fof(f10695, plain, (~ (e10 = op1(e13, e15)) | ~ spl42_8), inference(forward_demodulation, [], [f625, f1622])).
fof(f625, plain, ~ (op1(e13, e15) = op1(e16, e15)), inference(cnf_transformation, [], [f7])).
fof(f10668, plain, (~ spl42_135 | ~ spl42_100), inference(avatar_split_clause, [], [f10667, f2001, f2146])).
fof(f2146, plain, (spl42_135 <=> (e11 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_135])])).
fof(f10667, plain, (~ (e11 = op1(e14, e11)) | ~ spl42_100), inference(forward_demodulation, [], [f744, f2003])).
fof(f744, plain, ~ (op1(e14, e11) = op1(e14, e16)), inference(cnf_transformation, [], [f7])).
fof(f10662, plain, (~ spl42_128 | ~ spl42_100), inference(avatar_split_clause, [], [f10661, f2001, f2117])).
fof(f2117, plain, (spl42_128 <=> (e11 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_128])])).
fof(f10661, plain, (~ (e11 = op1(e14, e12)) | ~ spl42_100), inference(forward_demodulation, [], [f748, f2003])).
fof(f748, plain, ~ (op1(e14, e12) = op1(e14, e16)), inference(cnf_transformation, [], [f7])).
fof(f10656, plain, (~ spl42_121 | ~ spl42_100), inference(avatar_split_clause, [], [f10655, f2001, f2088])).
fof(f10655, plain, (~ (e11 = op1(e14, e13)) | ~ spl42_100), inference(forward_demodulation, [], [f751, f2003])).
fof(f751, plain, ~ (op1(e14, e13) = op1(e14, e16)), inference(cnf_transformation, [], [f7])).
fof(f10654, plain, (~ spl42_106 | ~ spl42_8), inference(avatar_split_clause, [], [f10653, f1620, f2026])).
fof(f2026, plain, (spl42_106 <=> (e10 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_106])])).
fof(f10653, plain, (~ (e10 = op1(e14, e15)) | ~ spl42_8), inference(forward_demodulation, [], [f627, f1622])).
fof(f627, plain, ~ (op1(e14, e15) = op1(e16, e15)), inference(cnf_transformation, [], [f7])).
fof(f10633, plain, (~ spl42_85 | ~ spl42_50), inference(avatar_split_clause, [], [f10632, f1794, f1939])).
fof(f1939, plain, (spl42_85 <=> (e10 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_85])])).
fof(f10632, plain, (~ (e10 = op1(e15, e11)) | ~ spl42_50), inference(forward_demodulation, [], [f765, f1796])).
fof(f765, plain, ~ (op1(e15, e11) = op1(e15, e16)), inference(cnf_transformation, [], [f7])).
fof(f10625, plain, (~ spl42_71 | ~ spl42_50), inference(avatar_split_clause, [], [f10624, f1794, f1881])).
fof(f1881, plain, (spl42_71 <=> (e10 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_71])])).
fof(f10624, plain, (~ (e10 = op1(e15, e13)) | ~ spl42_50), inference(forward_demodulation, [], [f772, f1796])).
fof(f772, plain, ~ (op1(e15, e13) = op1(e15, e16)), inference(cnf_transformation, [], [f7])).
fof(f10621, plain, (~ spl42_64 | ~ spl42_50), inference(avatar_split_clause, [], [f10620, f1794, f1852])).
fof(f1852, plain, (spl42_64 <=> (e10 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_64])])).
fof(f10620, plain, (~ (e10 = op1(e15, e14)) | ~ spl42_50), inference(forward_demodulation, [], [f774, f1796])).
fof(f774, plain, ~ (op1(e15, e14) = op1(e15, e16)), inference(cnf_transformation, [], [f7])).
fof(f10615, plain, (spl42_58 | ~ spl42_6), inference(avatar_split_clause, [], [f10512, f1611, f1827])).
fof(f10512, plain, ((e11 = op1(e15, e15)) | ~ spl42_6), inference(backward_demodulation, [], [f1187, f1613])).
fof(f10569, plain, (~ spl42_29 | ~ spl42_8), inference(avatar_split_clause, [], [f10568, f1620, f1707])).
fof(f10568, plain, (~ (e10 = op1(e16, e12)) | ~ spl42_8), inference(forward_demodulation, [], [f789, f1622])).
fof(f789, plain, ~ (op1(e16, e12) = op1(e16, e15)), inference(cnf_transformation, [], [f7])).
fof(f10548, plain, (~ spl42_24 | ~ spl42_26), inference(avatar_contradiction_clause, [], [f10547])).
fof(f10547, plain, ($false | (~ spl42_24 | ~ spl42_26)), inference(subsumption_resolution, [], [f10546, f1103])).
fof(f10546, plain, ((e12 = e14) | (~ spl42_24 | ~ spl42_26)), inference(forward_demodulation, [], [f1696, f1688])).
fof(f1688, plain, ((e12 = op1(e16, e13)) | ~ spl42_24), inference(avatar_component_clause, [], [f1686])).
fof(f1686, plain, (spl42_24 <=> (e12 = op1(e16, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_24])])).
fof(f10540, plain, (~ spl42_15 | ~ spl42_8), inference(avatar_split_clause, [], [f10539, f1620, f1649])).
fof(f1649, plain, (spl42_15 <=> (e10 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_15])])).
fof(f10539, plain, (~ (e10 = op1(e16, e14)) | ~ spl42_8), inference(forward_demodulation, [], [f794, f1622])).
fof(f794, plain, ~ (op1(e16, e14) = op1(e16, e15)), inference(cnf_transformation, [], [f7])).
fof(f10521, plain, (~ spl42_6 | ~ spl42_253), inference(avatar_contradiction_clause, [], [f10520])).
fof(f10520, plain, ($false | (~ spl42_6 | ~ spl42_253)), inference(subsumption_resolution, [], [f10519, f1092])).
fof(f1092, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f9])).
fof(f10519, plain, ((e10 = e12) | (~ spl42_6 | ~ spl42_253)), inference(forward_demodulation, [], [f10514, f2637])).
fof(f2637, plain, ((e10 = op1(e11, e15)) | ~ spl42_253), inference(avatar_component_clause, [], [f2635])).
fof(f10518, plain, (~ spl42_6 | ~ spl42_57), inference(avatar_contradiction_clause, [], [f10517])).
fof(f10517, plain, ($false | (~ spl42_6 | ~ spl42_57)), inference(subsumption_resolution, [], [f10516, f1091])).
fof(f10516, plain, ((e10 = e11) | (~ spl42_6 | ~ spl42_57)), inference(forward_demodulation, [], [f10512, f1825])).
fof(f1825, plain, ((e10 = op1(e15, e15)) | ~ spl42_57), inference(avatar_component_clause, [], [f1823])).
fof(f10404, plain, (~ spl42_38 | ~ spl42_39), inference(avatar_contradiction_clause, [], [f10403])).
fof(f10403, plain, ($false | (~ spl42_38 | ~ spl42_39)), inference(subsumption_resolution, [], [f10402, f1102])).
fof(f10402, plain, ((e12 = e13) | (~ spl42_38 | ~ spl42_39)), inference(backward_demodulation, [], [f1750, f1746])).
fof(f1746, plain, ((e12 = op1(e16, e11)) | ~ spl42_38), inference(avatar_component_clause, [], [f1744])).
fof(f1744, plain, (spl42_38 <=> (e12 = op1(e16, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_38])])).
fof(f10401, plain, (~ spl42_39 | ~ spl42_40), inference(avatar_contradiction_clause, [], [f10400])).
fof(f10400, plain, ($false | (~ spl42_39 | ~ spl42_40)), inference(subsumption_resolution, [], [f10399, f1106])).
fof(f10399, plain, ((e13 = e14) | (~ spl42_39 | ~ spl42_40)), inference(backward_demodulation, [], [f1754, f1750])).
fof(f1754, plain, ((e14 = op1(e16, e11)) | ~ spl42_40), inference(avatar_component_clause, [], [f1752])).
fof(f1752, plain, (spl42_40 <=> (e14 = op1(e16, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_40])])).
fof(f10242, plain, (~ spl42_100 | ~ spl42_101), inference(avatar_contradiction_clause, [], [f10241])).
fof(f10241, plain, ($false | (~ spl42_100 | ~ spl42_101)), inference(subsumption_resolution, [], [f10240, f1097])).
fof(f10240, plain, ((e11 = e12) | (~ spl42_100 | ~ spl42_101)), inference(backward_demodulation, [], [f2007, f2003])).
fof(f2007, plain, ((e12 = op1(e14, e16)) | ~ spl42_101), inference(avatar_component_clause, [], [f2005])).
fof(f9997, plain, (~ spl42_199 | ~ spl42_200), inference(avatar_contradiction_clause, [], [f9996])).
fof(f9996, plain, ($false | (~ spl42_199 | ~ spl42_200)), inference(subsumption_resolution, [], [f9995, f1102])).
fof(f9995, plain, ((e12 = e13) | (~ spl42_199 | ~ spl42_200)), inference(backward_demodulation, [], [f2417, f2413])).
fof(f2417, plain, ((e13 = op1(e12, e16)) | ~ spl42_200), inference(avatar_component_clause, [], [f2415])).
fof(f2415, plain, (spl42_200 <=> (e13 = op1(e12, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_200])])).
fof(f9894, plain, (~ spl42_134 | ~ spl42_232), inference(avatar_split_clause, [], [f9881, f2548, f2142])).
fof(f9881, plain, (~ (e10 = op1(e14, e11)) | ~ spl42_232), inference(backward_demodulation, [], [f536, f2550])).
fof(f536, plain, ~ (op1(e12, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f9785, plain, (~ spl42_267 | ~ spl42_330), inference(avatar_contradiction_clause, [], [f9784])).
fof(f9784, plain, ($false | (~ spl42_267 | ~ spl42_330)), inference(subsumption_resolution, [], [f9783, f1093])).
fof(f9783, plain, ((e10 = e13) | (~ spl42_267 | ~ spl42_330)), inference(forward_demodulation, [], [f9782, f2956])).
fof(f9782, plain, ((e13 = op1(e10, e11)) | (~ spl42_267 | ~ spl42_330)), inference(forward_demodulation, [], [f9773, f2956])).
fof(f9773, plain, ((e13 = op1(op1(e10, e11), e11)) | ~ spl42_267), inference(backward_demodulation, [], [f280, f2695])).
fof(f9763, plain, (~ spl42_274 | ~ spl42_330), inference(avatar_contradiction_clause, [], [f9762])).
fof(f9762, plain, ($false | (~ spl42_274 | ~ spl42_330)), inference(subsumption_resolution, [], [f9761, f1092])).
fof(f9761, plain, ((e10 = e12) | (~ spl42_274 | ~ spl42_330)), inference(forward_demodulation, [], [f9760, f2956])).
fof(f9760, plain, ((e12 = op1(e10, e11)) | (~ spl42_274 | ~ spl42_330)), inference(forward_demodulation, [], [f9750, f2956])).
fof(f9750, plain, ((e12 = op1(op1(e10, e11), e11)) | ~ spl42_274), inference(backward_demodulation, [], [f273, f2724])).
fof(f2724, plain, ((e10 = op1(e11, e12)) | ~ spl42_274), inference(avatar_component_clause, [], [f2722])).
fof(f2722, plain, (spl42_274 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_274])])).
fof(f9682, plain, (~ spl42_154 | ~ spl42_301), inference(avatar_split_clause, [], [f9676, f2833, f2224])).
fof(f9676, plain, (~ (e16 = op1(e13, e16)) | ~ spl42_301), inference(backward_demodulation, [], [f631, f2835])).
fof(f631, plain, ~ (op1(e10, e16) = op1(e13, e16)), inference(cnf_transformation, [], [f7])).
fof(f9582, plain, (~ spl42_309 | ~ spl42_330), inference(avatar_split_clause, [], [f9567, f2954, f2867])).
fof(f2867, plain, (spl42_309 <=> (e10 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_309])])).
fof(f9567, plain, (~ (e10 = op1(e10, e14)) | ~ spl42_330), inference(backward_demodulation, [], [f658, f2956])).
fof(f658, plain, ~ (op1(e10, e11) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f9573, plain, (~ spl42_330 | ~ spl42_337), inference(avatar_contradiction_clause, [], [f9572])).
fof(f9572, plain, ($false | (~ spl42_330 | ~ spl42_337)), inference(subsumption_resolution, [], [f9571, f1091])).
fof(f9571, plain, ((e10 = e11) | (~ spl42_330 | ~ spl42_337)), inference(forward_demodulation, [], [f9570, f2985])).
fof(f2985, plain, ((e10 = op1(e10, e10)) | ~ spl42_337), inference(avatar_component_clause, [], [f2983])).
fof(f2983, plain, (spl42_337 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_337])])).
fof(f9570, plain, ((op1(e10, e10) = e11) | (~ spl42_330 | ~ spl42_337)), inference(forward_demodulation, [], [f9558, f2985])).
fof(f9558, plain, ((e11 = op1(op1(e10, e10), e10)) | ~ spl42_330), inference(backward_demodulation, [], [f265, f2956])).
fof(f9524, plain, (spl42_688 | ~ spl42_349), inference(avatar_split_clause, [], [f9523, f3130, f4816])).
fof(f9523, plain, ((e25 = h7(e15)) | ~ spl42_349), inference(backward_demodulation, [], [f1246, f3132])).
fof(f9480, plain, (~ spl42_972 | ~ spl42_364), inference(avatar_split_clause, [], [f9471, f3192, f6460])).
fof(f6460, plain, (spl42_972 <=> (e26 = h7(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_972])])).
fof(f9471, plain, (~ (e26 = h7(e13)) | ~ spl42_364), inference(backward_demodulation, [], [f4799, f3194])).
fof(f9457, plain, (~ spl42_369 | ~ spl42_370), inference(avatar_contradiction_clause, [], [f9456])).
fof(f9456, plain, ($false | (~ spl42_369 | ~ spl42_370)), inference(subsumption_resolution, [], [f9455, f1130])).
fof(f9455, plain, ((e24 = e25) | (~ spl42_369 | ~ spl42_370)), inference(backward_demodulation, [], [f3219, f3215])).
fof(f3219, plain, ((e25 = op2(e26, e23)) | ~ spl42_370), inference(avatar_component_clause, [], [f3217])).
fof(f3217, plain, (spl42_370 <=> (e25 = op2(e26, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_370])])).
fof(f9399, plain, (~ spl42_381 | ~ spl42_382), inference(avatar_contradiction_clause, [], [f9398])).
fof(f9398, plain, ($false | (~ spl42_381 | ~ spl42_382)), inference(subsumption_resolution, [], [f9397, f1123])).
fof(f9397, plain, ((e22 = e23) | (~ spl42_381 | ~ spl42_382)), inference(backward_demodulation, [], [f3269, f3265])).
fof(f3265, plain, ((e22 = op2(e26, e21)) | ~ spl42_381), inference(avatar_component_clause, [], [f3263])).
fof(f9396, plain, (~ spl42_382 | ~ spl42_383), inference(avatar_contradiction_clause, [], [f9395])).
fof(f9395, plain, ($false | (~ spl42_382 | ~ spl42_383)), inference(subsumption_resolution, [], [f9394, f1127])).
fof(f9394, plain, ((e23 = e24) | (~ spl42_382 | ~ spl42_383)), inference(backward_demodulation, [], [f3273, f3269])).
fof(f3273, plain, ((e24 = op2(e26, e21)) | ~ spl42_383), inference(avatar_component_clause, [], [f3271])).
fof(f3271, plain, (spl42_383 <=> (e24 = op2(e26, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_383])])).
fof(f9387, plain, (spl42_972 | ~ spl42_385), inference(avatar_split_clause, [], [f9385, f3279, f6460])).
fof(f3279, plain, (spl42_385 <=> (e26 = op2(e26, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_385])])).
fof(f9385, plain, ((e26 = h7(e13)) | ~ spl42_385), inference(backward_demodulation, [], [f4789, f3281])).
fof(f3281, plain, ((e26 = op2(e26, e21)) | ~ spl42_385), inference(avatar_component_clause, [], [f3279])).
fof(f9368, plain, (~ spl42_972 | ~ spl42_392), inference(avatar_split_clause, [], [f9356, f3308, f6460])).
fof(f9356, plain, (~ (e26 = h7(e13)) | ~ spl42_392), inference(backward_demodulation, [], [f4796, f3310])).
fof(f9329, plain, (spl42_740 | ~ spl42_400), inference(avatar_split_clause, [], [f9328, f3342, f5098])).
fof(f9328, plain, ((e20 = h6(e15)) | ~ spl42_400), inference(backward_demodulation, [], [f1239, f3344])).
fof(f3344, plain, ((e20 = op2(e25, e25)) | ~ spl42_400), inference(avatar_component_clause, [], [f3342])).
fof(f9324, plain, (~ spl42_393 | ~ spl42_407), inference(avatar_split_clause, [], [f9316, f3371, f3313])).
fof(f3371, plain, (spl42_407 <=> (e20 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_407])])).
fof(f9316, plain, (~ (e20 = op2(e25, e26)) | ~ spl42_407), inference(backward_demodulation, [], [f1068, f3373])).
fof(f3373, plain, ((e20 = op2(e25, e24)) | ~ spl42_407), inference(avatar_component_clause, [], [f3371])).
fof(f1068, plain, ~ (op2(e25, e24) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f9311, plain, (~ spl42_393 | ~ spl42_414), inference(avatar_split_clause, [], [f9302, f3400, f3313])).
fof(f3400, plain, (spl42_414 <=> (e20 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_414])])).
fof(f9302, plain, (~ (e20 = op2(e25, e26)) | ~ spl42_414), inference(backward_demodulation, [], [f1066, f3402])).
fof(f3402, plain, ((e20 = op2(e25, e23)) | ~ spl42_414), inference(avatar_component_clause, [], [f3400])).
fof(f1066, plain, ~ (op2(e25, e23) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f9246, plain, (~ spl42_443 | ~ spl42_444), inference(avatar_contradiction_clause, [], [f9245])).
fof(f9245, plain, ($false | (~ spl42_443 | ~ spl42_444)), inference(subsumption_resolution, [], [f9244, f1118])).
fof(f9244, plain, ((e21 = e22) | (~ spl42_443 | ~ spl42_444)), inference(backward_demodulation, [], [f3526, f3522])).
fof(f9221, plain, (~ spl42_351 | ~ spl42_449), inference(avatar_split_clause, [], [f9213, f3545, f3139])).
fof(f9213, plain, (~ (e20 = op2(e26, e25)) | ~ spl42_449), inference(backward_demodulation, [], [f921, f3547])).
fof(f3547, plain, ((e20 = op2(e24, e25)) | ~ spl42_449), inference(avatar_component_clause, [], [f3545])).
fof(f921, plain, ~ (op2(e24, e25) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f9107, plain, (~ spl42_351 | ~ spl42_498), inference(avatar_split_clause, [], [f9098, f3748, f3139])).
fof(f9098, plain, (~ (e20 = op2(e26, e25)) | ~ spl42_498), inference(backward_demodulation, [], [f919, f3750])).
fof(f3750, plain, ((e20 = op2(e23, e25)) | ~ spl42_498), inference(avatar_component_clause, [], [f3748])).
fof(f919, plain, ~ (op2(e23, e25) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f8983, plain, (~ spl42_351 | ~ spl42_547), inference(avatar_split_clause, [], [f8973, f3951, f3139])).
fof(f3951, plain, (spl42_547 <=> (e20 = op2(e22, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_547])])).
fof(f8973, plain, (~ (e20 = op2(e26, e25)) | ~ spl42_547), inference(backward_demodulation, [], [f916, f3953])).
fof(f3953, plain, ((e20 = op2(e22, e25)) | ~ spl42_547), inference(avatar_component_clause, [], [f3951])).
fof(f916, plain, ~ (op2(e22, e25) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f8962, plain, (~ spl42_554 | ~ spl42_666), inference(avatar_contradiction_clause, [], [f8961])).
fof(f8961, plain, ($false | (~ spl42_554 | ~ spl42_666)), inference(subsumption_resolution, [], [f8960, f1115])).
fof(f8960, plain, ((e20 = e24) | (~ spl42_554 | ~ spl42_666)), inference(forward_demodulation, [], [f8959, f4446])).
fof(f4446, plain, ((e20 = op2(e20, e22)) | ~ spl42_666), inference(avatar_component_clause, [], [f4444])).
fof(f4444, plain, (spl42_666 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_666])])).
fof(f8959, plain, ((e24 = op2(e20, e22)) | (~ spl42_554 | ~ spl42_666)), inference(forward_demodulation, [], [f8951, f4446])).
fof(f8951, plain, ((e24 = op2(op2(e20, e22), e22)) | ~ spl42_554), inference(backward_demodulation, [], [f484, f3982])).
fof(f484, plain, (e24 = op2(op2(op2(e22, e24), e22), e22)), inference(cnf_transformation, [], [f6])).
fof(f8942, plain, (~ spl42_561 | ~ spl42_666), inference(avatar_contradiction_clause, [], [f8941])).
fof(f8941, plain, ($false | (~ spl42_561 | ~ spl42_666)), inference(subsumption_resolution, [], [f8940, f1114])).
fof(f8940, plain, ((e20 = e23) | (~ spl42_561 | ~ spl42_666)), inference(forward_demodulation, [], [f8939, f4446])).
fof(f8939, plain, ((e23 = op2(e20, e22)) | (~ spl42_561 | ~ spl42_666)), inference(forward_demodulation, [], [f8930, f4446])).
fof(f8930, plain, ((e23 = op2(op2(e20, e22), e22)) | ~ spl42_561), inference(backward_demodulation, [], [f477, f4011])).
fof(f4011, plain, ((e20 = op2(e22, e23)) | ~ spl42_561), inference(avatar_component_clause, [], [f4009])).
fof(f4009, plain, (spl42_561 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_561])])).
fof(f8874, plain, (~ spl42_592 | ~ spl42_593), inference(avatar_contradiction_clause, [], [f8873])).
fof(f8873, plain, ($false | (~ spl42_592 | ~ spl42_593)), inference(subsumption_resolution, [], [f8872, f1127])).
fof(f8872, plain, ((e23 = e24) | (~ spl42_592 | ~ spl42_593)), inference(backward_demodulation, [], [f4143, f4139])).
fof(f8848, plain, (~ spl42_351 | ~ spl42_596), inference(avatar_split_clause, [], [f8837, f4154, f3139])).
fof(f8837, plain, (~ (e20 = op2(e26, e25)) | ~ spl42_596), inference(backward_demodulation, [], [f912, f4156])).
fof(f4156, plain, ((e20 = op2(e21, e25)) | ~ spl42_596), inference(avatar_component_clause, [], [f4154])).
fof(f912, plain, ~ (op2(e21, e25) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f8824, plain, (~ spl42_603 | ~ spl42_673), inference(avatar_contradiction_clause, [], [f8823])).
fof(f8823, plain, ($false | (~ spl42_603 | ~ spl42_673)), inference(subsumption_resolution, [], [f8822, f1115])).
fof(f8822, plain, ((e20 = e24) | (~ spl42_603 | ~ spl42_673)), inference(forward_demodulation, [], [f8821, f4475])).
fof(f8821, plain, ((e24 = op2(e20, e21)) | (~ spl42_603 | ~ spl42_673)), inference(forward_demodulation, [], [f8812, f4475])).
fof(f8812, plain, ((e24 = op2(op2(e20, e21), e21)) | ~ spl42_603), inference(backward_demodulation, [], [f483, f4185])).
fof(f8802, plain, (~ spl42_610 | ~ spl42_673), inference(avatar_contradiction_clause, [], [f8801])).
fof(f8801, plain, ($false | (~ spl42_610 | ~ spl42_673)), inference(subsumption_resolution, [], [f8800, f1114])).
fof(f8800, plain, ((e20 = e23) | (~ spl42_610 | ~ spl42_673)), inference(forward_demodulation, [], [f8799, f4475])).
fof(f8799, plain, ((e23 = op2(e20, e21)) | (~ spl42_610 | ~ spl42_673)), inference(forward_demodulation, [], [f8789, f4475])).
fof(f8789, plain, ((e23 = op2(op2(e20, e21), e21)) | ~ spl42_610), inference(backward_demodulation, [], [f476, f4214])).
fof(f8778, plain, (~ spl42_617 | ~ spl42_673), inference(avatar_contradiction_clause, [], [f8777])).
fof(f8777, plain, ($false | (~ spl42_617 | ~ spl42_673)), inference(subsumption_resolution, [], [f8776, f1113])).
fof(f1113, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f10])).
fof(f8776, plain, ((e20 = e22) | (~ spl42_617 | ~ spl42_673)), inference(forward_demodulation, [], [f8775, f4475])).
fof(f8775, plain, ((e22 = op2(e20, e21)) | (~ spl42_617 | ~ spl42_673)), inference(forward_demodulation, [], [f8764, f4475])).
fof(f8751, plain, (~ spl42_582 | ~ spl42_631), inference(avatar_split_clause, [], [f8739, f4299, f4096])).
fof(f4096, plain, (spl42_582 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_582])])).
fof(f8739, plain, (~ (e20 = op2(e22, e20)) | ~ spl42_631), inference(backward_demodulation, [], [f803, f4301])).
fof(f803, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f8])).
fof(f8718, plain, (~ spl42_546 | ~ spl42_644), inference(avatar_split_clause, [], [f8712, f4352, f3946])).
fof(f3946, plain, (spl42_546 <=> (e26 = op2(e22, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_546])])).
fof(f8712, plain, (~ (e26 = op2(e22, e26)) | ~ spl42_644), inference(backward_demodulation, [], [f924, f4354])).
fof(f924, plain, ~ (op2(e20, e26) = op2(e22, e26)), inference(cnf_transformation, [], [f8])).
fof(f8717, plain, (~ spl42_595 | ~ spl42_644), inference(avatar_split_clause, [], [f8711, f4352, f4149])).
fof(f4149, plain, (spl42_595 <=> (e26 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_595])])).
fof(f8711, plain, (~ (e26 = op2(e21, e26)) | ~ spl42_644), inference(backward_demodulation, [], [f923, f4354])).
fof(f923, plain, ~ (op2(e20, e26) = op2(e21, e26)), inference(cnf_transformation, [], [f8])).
fof(f8707, plain, (~ spl42_351 | ~ spl42_645), inference(avatar_split_clause, [], [f8696, f4357, f3139])).
fof(f8696, plain, (~ (e20 = op2(e26, e25)) | ~ spl42_645), inference(backward_demodulation, [], [f907, f4359])).
fof(f4359, plain, ((e20 = op2(e20, e25)) | ~ spl42_645), inference(avatar_component_clause, [], [f4357])).
fof(f907, plain, ~ (op2(e20, e25) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f8615, plain, (~ spl42_666 | ~ spl42_673), inference(avatar_split_clause, [], [f8600, f4473, f4444])).
fof(f8600, plain, (~ (e20 = op2(e20, e22)) | ~ spl42_673), inference(backward_demodulation, [], [f950, f4475])).
fof(f8612, plain, (~ spl42_526 | ~ spl42_673), inference(avatar_split_clause, [], [f8597, f4473, f3864])).
fof(f3864, plain, (spl42_526 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_526])])).
fof(f8597, plain, (~ (e20 = op2(e23, e21)) | ~ spl42_673), inference(backward_demodulation, [], [f820, f4475])).
fof(f820, plain, ~ (op2(e20, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f8])).
fof(f8610, plain, (~ spl42_673 | ~ spl42_680), inference(avatar_contradiction_clause, [], [f8609])).
fof(f8609, plain, ($false | (~ spl42_673 | ~ spl42_680)), inference(subsumption_resolution, [], [f8608, f1112])).
fof(f8608, plain, ((e20 = e21) | (~ spl42_673 | ~ spl42_680)), inference(forward_demodulation, [], [f8607, f4504])).
fof(f4504, plain, ((e20 = op2(e20, e20)) | ~ spl42_680), inference(avatar_component_clause, [], [f4502])).
fof(f4502, plain, (spl42_680 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_680])])).
fof(f8607, plain, ((op2(e20, e20) = e21) | (~ spl42_673 | ~ spl42_680)), inference(forward_demodulation, [], [f8595, f4504])).
fof(f8595, plain, ((e21 = op2(op2(e20, e20), e20)) | ~ spl42_673), inference(backward_demodulation, [], [f461, f4475])).
fof(f8081, plain, (~ spl42_1193 | ~ spl42_1194 | ~ spl42_1195 | ~ spl42_1196 | ~ spl42_1197 | ~ spl42_1198 | ~ spl42_1199 | ~ spl42_1200 | ~ spl42_1201 | ~ spl42_1202 | ~ spl42_1203 | ~ spl42_1204 | ~ spl42_1205 | ~ spl42_1206 | ~ spl42_1207 | ~ spl42_1208 | ~ spl42_1209 | ~ spl42_1210 | ~ spl42_1211 | ~ spl42_1212 | ~ spl42_1213 | ~ spl42_1214 | ~ spl42_1215 | ~ spl42_1216 | ~ spl42_1217 | ~ spl42_1218 | ~ spl42_1219 | ~ spl42_1220 | ~ spl42_1221 | ~ spl42_1222 | ~ spl42_1223 | ~ spl42_1224 | ~ spl42_1225 | ~ spl42_1226 | spl42_879 | spl42_872 | spl42_865 | spl42_858 | spl42_851 | ~ spl42_1227 | ~ spl42_1228 | ~ spl42_1229 | ~ spl42_1230 | ~ spl42_1231 | ~ spl42_1232 | ~ spl42_1233 | ~ spl42_1234 | ~ spl42_1235 | ~ spl42_1236 | ~ spl42_1237 | ~ spl42_1238 | ~ spl42_1239 | ~ spl42_1240 | ~ spl42_1241 | ~ spl42_1242), inference(avatar_split_clause, [], [f7880, f8078, f8074, f8070, f8066, f8062, f8058, f8054, f8050, f8046, f8042, f8038, f8034, f8030, f8026, f8022, f8018, f5657, f5692, f5727, f5762, f5798, f8014, f8010, f8006, f8002, f7998, f7994, f7990, f7986, f7982, f7978, f7974, f7970, f7966, f7962, f7958, f7954, f7950, f7946, f7942, f7938, f7934, f7930, f7926, f7922, f7918, f7914, f7910, f7906, f7902, f7898, f7894, f7890, f7886, f7882])).
fof(f5798, plain, (spl42_879 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl42_879])])).
fof(f5762, plain, (spl42_872 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl42_872])])).
fof(f5727, plain, (spl42_865 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl42_865])])).
fof(f5692, plain, (spl42_858 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl42_858])])).
fof(f5657, plain, (spl42_851 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl42_851])])).
fof(f7880, plain, (~ (h2(op1(e10, e16)) = op2(h2(e10), e21)) | ~ (h2(e12) = h2(op1(e11, e15))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), e21)) | ~ (h2(op1(e12, e16)) = op2(h2(e12), e21)) | ~ (h2(op1(e13, e16)) = op2(h2(e13), e21)) | ~ (h2(op1(e14, e16)) = op2(h2(e14), e21)) | ~ (h2(e11) = h2(op1(e15, e15))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), e21)) | ~ (h2(op1(e16, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e16, e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7879, f1205])).
fof(f7879, plain, (~ (h2(e12) = h2(op1(e11, e15))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), e21)) | ~ (h2(op1(e12, e16)) = op2(h2(e12), e21)) | ~ (h2(op1(e13, e16)) = op2(h2(e13), e21)) | ~ (h2(op1(e14, e16)) = op2(h2(e14), e21)) | ~ (h2(e11) = h2(op1(e15, e15))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), e21)) | ~ (h2(op1(e16, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e16, e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7878, f4683])).
fof(f7878, plain, (~ (h2(op1(e11, e16)) = op2(h2(e11), e21)) | ~ (h2(op1(e12, e16)) = op2(h2(e12), e21)) | ~ (h2(op1(e13, e16)) = op2(h2(e13), e21)) | ~ (h2(op1(e14, e16)) = op2(h2(e14), e21)) | ~ (h2(e11) = h2(op1(e15, e15))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), e21)) | ~ (h2(op1(e16, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e16, e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7877, f1205])).
fof(f7877, plain, (~ (h2(op1(e12, e16)) = op2(h2(e12), e21)) | ~ (h2(op1(e13, e16)) = op2(h2(e13), e21)) | ~ (h2(op1(e14, e16)) = op2(h2(e14), e21)) | ~ (h2(e11) = h2(op1(e15, e15))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), e21)) | ~ (h2(op1(e16, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e16, e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7876, f1205])).
fof(f7876, plain, (~ (h2(op1(e13, e16)) = op2(h2(e13), e21)) | ~ (h2(op1(e14, e16)) = op2(h2(e14), e21)) | ~ (h2(e11) = h2(op1(e15, e15))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), e21)) | ~ (h2(op1(e16, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e16, e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7875, f1205])).
fof(f7875, plain, (~ (h2(op1(e14, e16)) = op2(h2(e14), e21)) | ~ (h2(e11) = h2(op1(e15, e15))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), e21)) | ~ (h2(op1(e16, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e16, e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7874, f1205])).
fof(f7874, plain, (~ (h2(e11) = h2(op1(e15, e15))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), e21)) | ~ (h2(op1(e16, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e16, e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7873, f4682])).
fof(f7873, plain, (~ (h2(op1(e15, e16)) = op2(h2(e15), e21)) | ~ (h2(op1(e16, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e16, e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7872, f1205])).
fof(f7872, plain, (~ (h2(op1(e16, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e16, e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7871, f1205])).
fof(f7871, plain, (~ (h2(e13) = h2(op1(e16, e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7870, f4684])).
fof(f7870, plain, (~ (h2(op1(e16, e11)) = op2(e21, h2(e11))) | ~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7869, f1205])).
fof(f7869, plain, (~ (h2(op1(e16, e12)) = op2(e21, h2(e12))) | ~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7868, f1205])).
fof(f7868, plain, (~ (h2(e14) = h2(op1(e16, e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7867, f4680])).
fof(f7867, plain, (~ (h2(op1(e16, e13)) = op2(e21, h2(e13))) | ~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7866, f1205])).
fof(f7866, plain, (~ (h2(op1(e16, e14)) = op2(e21, h2(e14))) | ~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7865, f1205])).
fof(f7865, plain, (~ (h2(e10) = h2(op1(e16, e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7864, f4685])).
fof(f7864, plain, (~ (h2(op1(e16, e15)) = op2(e21, h2(e15))) | ~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7863, f1205])).
fof(f7863, plain, (~ (h2(e15) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) | ~ (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7862, f1211])).
fof(f7862, plain, (~ (op2(e21, e21) = h2(op1(e16, e16))) | ~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) | ~ (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f7861, f1205])).
fof(f7861, plain, (~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP6 | ~ (h2(op1(e16, e16)) = op2(h2(e16), h2(e16))) | ~ (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) | ~ (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(subsumption_resolution, [], [f1553, f5795])).
fof(f5795, plain, ~ sP7, inference(subsumption_resolution, [], [f1491, f1205])).
fof(f1491, plain, (~ (e21 = h2(e16)) | ~ sP7), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e21 = h2(e16)) & ~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP7), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ((~ (e21 = h2(e16)) & ~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP7), inference(usedef, [], [e33])).
fof(e33, plain, (sP7 <=> (~ (e21 = h2(e16)) & ~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f1553, plain, (~ (e26 = h2(e15)) | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | ~ (h2(op1(e16, e16)) = op2(h2(e16), h2(e16))) | ~ (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) | ~ (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(cnf_transformation, [], [f68])).
fof(f68, plain, (((~ (e26 = h7(e16)) & ~ (e26 = h7(e15)) & ~ (e26 = h7(e14)) & ~ (e26 = h7(e13)) & ~ (e26 = h7(e12)) & ~ (e26 = h7(e11)) & ~ (e26 = h7(e10))) | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | ~ (h7(op1(e16, e16)) = op2(h7(e16), h7(e16))) | ~ (h7(op1(e16, e15)) = op2(h7(e16), h7(e15))) | ~ (h7(op1(e16, e14)) = op2(h7(e16), h7(e14))) | ~ (h7(op1(e16, e13)) = op2(h7(e16), h7(e13))) | ~ (h7(op1(e16, e12)) = op2(h7(e16), h7(e12))) | ~ (h7(op1(e16, e11)) = op2(h7(e16), h7(e11))) | ~ (h7(op1(e16, e10)) = op2(h7(e16), h7(e10))) | ~ (h7(op1(e15, e16)) = op2(h7(e15), h7(e16))) | ~ (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) | ~ (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) | ~ (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) | ~ (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) | ~ (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) | ~ (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) | ~ (h7(op1(e14, e16)) = op2(h7(e14), h7(e16))) | ~ (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e16)) = op2(h7(e13), h7(e16))) | ~ (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e16)) = op2(h7(e12), h7(e16))) | ~ (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e16)) = op2(h7(e11), h7(e16))) | ~ (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e16)) = op2(h7(e10), h7(e16))) | ~ (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e26 = h6(e16)) & ~ (e26 = h6(e15)) & ~ (e26 = h6(e14)) & ~ (e26 = h6(e13)) & ~ (e26 = h6(e12)) & ~ (e26 = h6(e11)) & ~ (e26 = h6(e10))) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e15)) = op2(h6(e16), h6(e15))) | ~ (h6(op1(e16, e14)) = op2(h6(e16), h6(e14))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) | ~ (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e26 = h5(e16)) & ~ (e26 = h5(e15)) & ~ (e26 = h5(e14)) & ~ (e26 = h5(e13)) & ~ (e26 = h5(e12)) & ~ (e26 = h5(e11)) & ~ (e26 = h5(e10))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | ~ (h5(op1(e16, e16)) = op2(h5(e16), h5(e16))) | ~ (h5(op1(e16, e15)) = op2(h5(e16), h5(e15))) | ~ (h5(op1(e16, e14)) = op2(h5(e16), h5(e14))) | ~ (h5(op1(e16, e13)) = op2(h5(e16), h5(e13))) | ~ (h5(op1(e16, e12)) = op2(h5(e16), h5(e12))) | ~ (h5(op1(e16, e11)) = op2(h5(e16), h5(e11))) | ~ (h5(op1(e16, e10)) = op2(h5(e16), h5(e10))) | ~ (h5(op1(e15, e16)) = op2(h5(e15), h5(e16))) | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e16)) = op2(h5(e14), h5(e16))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e16)) = op2(h5(e13), h5(e16))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e16)) = op2(h5(e12), h5(e16))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e16)) = op2(h5(e11), h5(e16))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e16)) = op2(h5(e10), h5(e16))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e26 = h4(e16)) & ~ (e26 = h4(e15)) & ~ (e26 = h4(e14)) & ~ (e26 = h4(e13)) & ~ (e26 = h4(e12)) & ~ (e26 = h4(e11)) & ~ (e26 = h4(e10))) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | ~ (h4(op1(e16, e16)) = op2(h4(e16), h4(e16))) | ~ (h4(op1(e16, e15)) = op2(h4(e16), h4(e15))) | ~ (h4(op1(e16, e14)) = op2(h4(e16), h4(e14))) | ~ (h4(op1(e16, e13)) = op2(h4(e16), h4(e13))) | ~ (h4(op1(e16, e12)) = op2(h4(e16), h4(e12))) | ~ (h4(op1(e16, e11)) = op2(h4(e16), h4(e11))) | ~ (h4(op1(e16, e10)) = op2(h4(e16), h4(e10))) | ~ (h4(op1(e15, e16)) = op2(h4(e15), h4(e16))) | ~ (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) | ~ (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) | ~ (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) | ~ (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) | ~ (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) | ~ (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) | ~ (h4(op1(e14, e16)) = op2(h4(e14), h4(e16))) | ~ (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e16)) = op2(h4(e13), h4(e16))) | ~ (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e16)) = op2(h4(e12), h4(e16))) | ~ (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e16)) = op2(h4(e11), h4(e16))) | ~ (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e16)) = op2(h4(e10), h4(e16))) | ~ (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e26 = h3(e16)) & ~ (e26 = h3(e15)) & ~ (e26 = h3(e14)) & ~ (e26 = h3(e13)) & ~ (e26 = h3(e12)) & ~ (e26 = h3(e11)) & ~ (e26 = h3(e10))) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | ~ (h3(op1(e16, e16)) = op2(h3(e16), h3(e16))) | ~ (h3(op1(e16, e15)) = op2(h3(e16), h3(e15))) | ~ (h3(op1(e16, e14)) = op2(h3(e16), h3(e14))) | ~ (h3(op1(e16, e13)) = op2(h3(e16), h3(e13))) | ~ (h3(op1(e16, e12)) = op2(h3(e16), h3(e12))) | ~ (h3(op1(e16, e11)) = op2(h3(e16), h3(e11))) | ~ (h3(op1(e16, e10)) = op2(h3(e16), h3(e10))) | ~ (h3(op1(e15, e16)) = op2(h3(e15), h3(e16))) | ~ (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) | ~ (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) | ~ (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) | ~ (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) | ~ (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) | ~ (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) | ~ (h3(op1(e14, e16)) = op2(h3(e14), h3(e16))) | ~ (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e16)) = op2(h3(e13), h3(e16))) | ~ (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e16)) = op2(h3(e12), h3(e16))) | ~ (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e16)) = op2(h3(e11), h3(e16))) | ~ (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e16)) = op2(h3(e10), h3(e16))) | ~ (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e26 = h2(e16)) & ~ (e26 = h2(e15)) & ~ (e26 = h2(e14)) & ~ (e26 = h2(e13)) & ~ (e26 = h2(e12)) & ~ (e26 = h2(e11)) & ~ (e26 = h2(e10))) | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | ~ (h2(op1(e16, e16)) = op2(h2(e16), h2(e16))) | ~ (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) | ~ (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e26 = h1(e16)) & ~ (e26 = h1(e15)) & ~ (e26 = h1(e14)) & ~ (e26 = h1(e13)) & ~ (e26 = h1(e12)) & ~ (e26 = h1(e11)) & ~ (e26 = h1(e10))) | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (h1(op1(e16, e16)) = op2(h1(e16), h1(e16))) | ~ (h1(op1(e16, e15)) = op2(h1(e16), h1(e15))) | ~ (h1(op1(e16, e14)) = op2(h1(e16), h1(e14))) | ~ (h1(op1(e16, e13)) = op2(h1(e16), h1(e13))) | ~ (h1(op1(e16, e12)) = op2(h1(e16), h1(e12))) | ~ (h1(op1(e16, e11)) = op2(h1(e16), h1(e11))) | ~ (h1(op1(e16, e10)) = op2(h1(e16), h1(e10))) | ~ (h1(op1(e15, e16)) = op2(h1(e15), h1(e16))) | ~ (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) | ~ (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) | ~ (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) | ~ (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) | ~ (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) | ~ (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) | ~ (h1(op1(e14, e16)) = op2(h1(e14), h1(e16))) | ~ (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e16)) = op2(h1(e13), h1(e16))) | ~ (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e16)) = op2(h1(e12), h1(e16))) | ~ (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e16)) = op2(h1(e11), h1(e16))) | ~ (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e16)) = op2(h1(e10), h1(e16))) | ~ (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f25, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26])).
fof(f26, plain, ((~ (e20 = h1(e16)) & ~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(usedef, [], [e26])).
fof(e26, plain, (sP0 <=> (~ (e20 = h1(e16)) & ~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f27, plain, ((~ (e21 = h1(e16)) & ~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(usedef, [], [e27])).
fof(e27, plain, (sP1 <=> (~ (e21 = h1(e16)) & ~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f28, plain, ((~ (e22 = h1(e16)) & ~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(usedef, [], [e28])).
fof(e28, plain, (sP2 <=> (~ (e22 = h1(e16)) & ~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f29, plain, ((~ (e23 = h1(e16)) & ~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | ~ sP3), inference(usedef, [], [e29])).
fof(e29, plain, (sP3 <=> (~ (e23 = h1(e16)) & ~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f30, plain, ((~ (e24 = h1(e16)) & ~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | ~ sP4), inference(usedef, [], [e30])).
fof(e30, plain, (sP4 <=> (~ (e24 = h1(e16)) & ~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f31, plain, ((~ (e25 = h1(e16)) & ~ (e25 = h1(e15)) & ~ (e25 = h1(e14)) & ~ (e25 = h1(e13)) & ~ (e25 = h1(e12)) & ~ (e25 = h1(e11)) & ~ (e25 = h1(e10))) | ~ sP5), inference(usedef, [], [e31])).
fof(e31, plain, (sP5 <=> (~ (e25 = h1(e16)) & ~ (e25 = h1(e15)) & ~ (e25 = h1(e14)) & ~ (e25 = h1(e13)) & ~ (e25 = h1(e12)) & ~ (e25 = h1(e11)) & ~ (e25 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f32, plain, ((~ (e20 = h2(e16)) & ~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP6), inference(usedef, [], [e32])).
fof(e32, plain, (sP6 <=> (~ (e20 = h2(e16)) & ~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f34, plain, ((~ (e22 = h2(e16)) & ~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP8), inference(usedef, [], [e34])).
fof(e34, plain, (sP8 <=> (~ (e22 = h2(e16)) & ~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f35, plain, ((~ (e23 = h2(e16)) & ~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP9), inference(usedef, [], [e35])).
fof(e35, plain, (sP9 <=> (~ (e23 = h2(e16)) & ~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f36, plain, ((~ (e24 = h2(e16)) & ~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | ~ sP10), inference(usedef, [], [e36])).
fof(e36, plain, (sP10 <=> (~ (e24 = h2(e16)) & ~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f37, plain, ((~ (e25 = h2(e16)) & ~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10))) | ~ sP11), inference(usedef, [], [e37])).
fof(e37, plain, (sP11 <=> (~ (e25 = h2(e16)) & ~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f38, plain, ((~ (e20 = h3(e16)) & ~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP12), inference(usedef, [], [e38])).
fof(e38, plain, (sP12 <=> (~ (e20 = h3(e16)) & ~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f39, plain, ((~ (e21 = h3(e16)) & ~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP13), inference(usedef, [], [e39])).
fof(e39, plain, (sP13 <=> (~ (e21 = h3(e16)) & ~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f40, plain, ((~ (e22 = h3(e16)) & ~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP14), inference(usedef, [], [e40])).
fof(e40, plain, (sP14 <=> (~ (e22 = h3(e16)) & ~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f41, plain, ((~ (e23 = h3(e16)) & ~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP15), inference(usedef, [], [e41])).
fof(e41, plain, (sP15 <=> (~ (e23 = h3(e16)) & ~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f42, plain, ((~ (e24 = h3(e16)) & ~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | ~ sP16), inference(usedef, [], [e42])).
fof(e42, plain, (sP16 <=> (~ (e24 = h3(e16)) & ~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f43, plain, ((~ (e25 = h3(e16)) & ~ (e25 = h3(e15)) & ~ (e25 = h3(e14)) & ~ (e25 = h3(e13)) & ~ (e25 = h3(e12)) & ~ (e25 = h3(e11)) & ~ (e25 = h3(e10))) | ~ sP17), inference(usedef, [], [e43])).
fof(e43, plain, (sP17 <=> (~ (e25 = h3(e16)) & ~ (e25 = h3(e15)) & ~ (e25 = h3(e14)) & ~ (e25 = h3(e13)) & ~ (e25 = h3(e12)) & ~ (e25 = h3(e11)) & ~ (e25 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f44, plain, ((~ (e20 = h4(e16)) & ~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP18), inference(usedef, [], [e44])).
fof(e44, plain, (sP18 <=> (~ (e20 = h4(e16)) & ~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f45, plain, ((~ (e21 = h4(e16)) & ~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP19), inference(usedef, [], [e45])).
fof(e45, plain, (sP19 <=> (~ (e21 = h4(e16)) & ~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f46, plain, ((~ (e22 = h4(e16)) & ~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP20), inference(usedef, [], [e46])).
fof(e46, plain, (sP20 <=> (~ (e22 = h4(e16)) & ~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f47, plain, ((~ (e23 = h4(e16)) & ~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | ~ sP21), inference(usedef, [], [e47])).
fof(e47, plain, (sP21 <=> (~ (e23 = h4(e16)) & ~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f48, plain, ((~ (e24 = h4(e16)) & ~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | ~ sP22), inference(usedef, [], [e48])).
fof(e48, plain, (sP22 <=> (~ (e24 = h4(e16)) & ~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f49, plain, ((~ (e25 = h4(e16)) & ~ (e25 = h4(e15)) & ~ (e25 = h4(e14)) & ~ (e25 = h4(e13)) & ~ (e25 = h4(e12)) & ~ (e25 = h4(e11)) & ~ (e25 = h4(e10))) | ~ sP23), inference(usedef, [], [e49])).
fof(e49, plain, (sP23 <=> (~ (e25 = h4(e16)) & ~ (e25 = h4(e15)) & ~ (e25 = h4(e14)) & ~ (e25 = h4(e13)) & ~ (e25 = h4(e12)) & ~ (e25 = h4(e11)) & ~ (e25 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f50, plain, ((~ (e20 = h5(e16)) & ~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP24), inference(usedef, [], [e50])).
fof(e50, plain, (sP24 <=> (~ (e20 = h5(e16)) & ~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f51, plain, ((~ (e21 = h5(e16)) & ~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP25), inference(usedef, [], [e51])).
fof(e51, plain, (sP25 <=> (~ (e21 = h5(e16)) & ~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f52, plain, ((~ (e22 = h5(e16)) & ~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP26), inference(usedef, [], [e52])).
fof(e52, plain, (sP26 <=> (~ (e22 = h5(e16)) & ~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f53, plain, ((~ (e23 = h5(e16)) & ~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP27), inference(usedef, [], [e53])).
fof(e53, plain, (sP27 <=> (~ (e23 = h5(e16)) & ~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f54, plain, ((~ (e24 = h5(e16)) & ~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | ~ sP28), inference(usedef, [], [e54])).
fof(e54, plain, (sP28 <=> (~ (e24 = h5(e16)) & ~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f55, plain, ((~ (e25 = h5(e16)) & ~ (e25 = h5(e15)) & ~ (e25 = h5(e14)) & ~ (e25 = h5(e13)) & ~ (e25 = h5(e12)) & ~ (e25 = h5(e11)) & ~ (e25 = h5(e10))) | ~ sP29), inference(usedef, [], [e55])).
fof(e55, plain, (sP29 <=> (~ (e25 = h5(e16)) & ~ (e25 = h5(e15)) & ~ (e25 = h5(e14)) & ~ (e25 = h5(e13)) & ~ (e25 = h5(e12)) & ~ (e25 = h5(e11)) & ~ (e25 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f56, plain, ((~ (e20 = h6(e16)) & ~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP30), inference(usedef, [], [e56])).
fof(e56, plain, (sP30 <=> (~ (e20 = h6(e16)) & ~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f57, plain, ((~ (e21 = h6(e16)) & ~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP31), inference(usedef, [], [e57])).
fof(e57, plain, (sP31 <=> (~ (e21 = h6(e16)) & ~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f58, plain, ((~ (e22 = h6(e16)) & ~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP32), inference(usedef, [], [e58])).
fof(e58, plain, (sP32 <=> (~ (e22 = h6(e16)) & ~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f59, plain, ((~ (e23 = h6(e16)) & ~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | ~ sP33), inference(usedef, [], [e59])).
fof(e59, plain, (sP33 <=> (~ (e23 = h6(e16)) & ~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f60, plain, ((~ (e24 = h6(e16)) & ~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | ~ sP34), inference(usedef, [], [e60])).
fof(e60, plain, (sP34 <=> (~ (e24 = h6(e16)) & ~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f61, plain, ((~ (e25 = h6(e16)) & ~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10))) | ~ sP35), inference(usedef, [], [e61])).
fof(e61, plain, (sP35 <=> (~ (e25 = h6(e16)) & ~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f62, plain, ((~ (e20 = h7(e16)) & ~ (e20 = h7(e15)) & ~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP36), inference(usedef, [], [e62])).
fof(e62, plain, (sP36 <=> (~ (e20 = h7(e16)) & ~ (e20 = h7(e15)) & ~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f63, plain, ((~ (e21 = h7(e16)) & ~ (e21 = h7(e15)) & ~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP37), inference(usedef, [], [e63])).
fof(e63, plain, (sP37 <=> (~ (e21 = h7(e16)) & ~ (e21 = h7(e15)) & ~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f64, plain, ((~ (e22 = h7(e16)) & ~ (e22 = h7(e15)) & ~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP38), inference(usedef, [], [e64])).
fof(e64, plain, (sP38 <=> (~ (e22 = h7(e16)) & ~ (e22 = h7(e15)) & ~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f65, plain, ((~ (e23 = h7(e16)) & ~ (e23 = h7(e15)) & ~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | ~ sP39), inference(usedef, [], [e65])).
fof(e65, plain, (sP39 <=> (~ (e23 = h7(e16)) & ~ (e23 = h7(e15)) & ~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f66, plain, ((~ (e24 = h7(e16)) & ~ (e24 = h7(e15)) & ~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | ~ sP40), inference(usedef, [], [e66])).
fof(e66, plain, (sP40 <=> (~ (e24 = h7(e16)) & ~ (e24 = h7(e15)) & ~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f67, plain, ((~ (e25 = h7(e16)) & ~ (e25 = h7(e15)) & ~ (e25 = h7(e14)) & ~ (e25 = h7(e13)) & ~ (e25 = h7(e12)) & ~ (e25 = h7(e11)) & ~ (e25 = h7(e10))) | ~ sP41), inference(usedef, [], [e67])).
fof(e67, plain, (sP41 <=> (~ (e25 = h7(e16)) & ~ (e25 = h7(e15)) & ~ (e25 = h7(e14)) & ~ (e25 = h7(e13)) & ~ (e25 = h7(e12)) & ~ (e25 = h7(e11)) & ~ (e25 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f25, plain, (((~ (e26 = h7(e16)) & ~ (e26 = h7(e15)) & ~ (e26 = h7(e14)) & ~ (e26 = h7(e13)) & ~ (e26 = h7(e12)) & ~ (e26 = h7(e11)) & ~ (e26 = h7(e10))) | (~ (e25 = h7(e16)) & ~ (e25 = h7(e15)) & ~ (e25 = h7(e14)) & ~ (e25 = h7(e13)) & ~ (e25 = h7(e12)) & ~ (e25 = h7(e11)) & ~ (e25 = h7(e10))) | (~ (e24 = h7(e16)) & ~ (e24 = h7(e15)) & ~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | (~ (e23 = h7(e16)) & ~ (e23 = h7(e15)) & ~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | (~ (e22 = h7(e16)) & ~ (e22 = h7(e15)) & ~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | (~ (e21 = h7(e16)) & ~ (e21 = h7(e15)) & ~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | (~ (e20 = h7(e16)) & ~ (e20 = h7(e15)) & ~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ (h7(op1(e16, e16)) = op2(h7(e16), h7(e16))) | ~ (h7(op1(e16, e15)) = op2(h7(e16), h7(e15))) | ~ (h7(op1(e16, e14)) = op2(h7(e16), h7(e14))) | ~ (h7(op1(e16, e13)) = op2(h7(e16), h7(e13))) | ~ (h7(op1(e16, e12)) = op2(h7(e16), h7(e12))) | ~ (h7(op1(e16, e11)) = op2(h7(e16), h7(e11))) | ~ (h7(op1(e16, e10)) = op2(h7(e16), h7(e10))) | ~ (h7(op1(e15, e16)) = op2(h7(e15), h7(e16))) | ~ (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) | ~ (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) | ~ (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) | ~ (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) | ~ (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) | ~ (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) | ~ (h7(op1(e14, e16)) = op2(h7(e14), h7(e16))) | ~ (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e16)) = op2(h7(e13), h7(e16))) | ~ (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e16)) = op2(h7(e12), h7(e16))) | ~ (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e16)) = op2(h7(e11), h7(e16))) | ~ (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e16)) = op2(h7(e10), h7(e16))) | ~ (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e26 = h6(e16)) & ~ (e26 = h6(e15)) & ~ (e26 = h6(e14)) & ~ (e26 = h6(e13)) & ~ (e26 = h6(e12)) & ~ (e26 = h6(e11)) & ~ (e26 = h6(e10))) | (~ (e25 = h6(e16)) & ~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10))) | (~ (e24 = h6(e16)) & ~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | (~ (e23 = h6(e16)) & ~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e16)) & ~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e16)) & ~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e16)) & ~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e15)) = op2(h6(e16), h6(e15))) | ~ (h6(op1(e16, e14)) = op2(h6(e16), h6(e14))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) | ~ (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e26 = h5(e16)) & ~ (e26 = h5(e15)) & ~ (e26 = h5(e14)) & ~ (e26 = h5(e13)) & ~ (e26 = h5(e12)) & ~ (e26 = h5(e11)) & ~ (e26 = h5(e10))) | (~ (e25 = h5(e16)) & ~ (e25 = h5(e15)) & ~ (e25 = h5(e14)) & ~ (e25 = h5(e13)) & ~ (e25 = h5(e12)) & ~ (e25 = h5(e11)) & ~ (e25 = h5(e10))) | (~ (e24 = h5(e16)) & ~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | (~ (e23 = h5(e16)) & ~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e16)) & ~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e16)) & ~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e16)) & ~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e16, e16)) = op2(h5(e16), h5(e16))) | ~ (h5(op1(e16, e15)) = op2(h5(e16), h5(e15))) | ~ (h5(op1(e16, e14)) = op2(h5(e16), h5(e14))) | ~ (h5(op1(e16, e13)) = op2(h5(e16), h5(e13))) | ~ (h5(op1(e16, e12)) = op2(h5(e16), h5(e12))) | ~ (h5(op1(e16, e11)) = op2(h5(e16), h5(e11))) | ~ (h5(op1(e16, e10)) = op2(h5(e16), h5(e10))) | ~ (h5(op1(e15, e16)) = op2(h5(e15), h5(e16))) | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e16)) = op2(h5(e14), h5(e16))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e16)) = op2(h5(e13), h5(e16))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e16)) = op2(h5(e12), h5(e16))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e16)) = op2(h5(e11), h5(e16))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e16)) = op2(h5(e10), h5(e16))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e26 = h4(e16)) & ~ (e26 = h4(e15)) & ~ (e26 = h4(e14)) & ~ (e26 = h4(e13)) & ~ (e26 = h4(e12)) & ~ (e26 = h4(e11)) & ~ (e26 = h4(e10))) | (~ (e25 = h4(e16)) & ~ (e25 = h4(e15)) & ~ (e25 = h4(e14)) & ~ (e25 = h4(e13)) & ~ (e25 = h4(e12)) & ~ (e25 = h4(e11)) & ~ (e25 = h4(e10))) | (~ (e24 = h4(e16)) & ~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | (~ (e23 = h4(e16)) & ~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e16)) & ~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e16)) & ~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e16)) & ~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e16, e16)) = op2(h4(e16), h4(e16))) | ~ (h4(op1(e16, e15)) = op2(h4(e16), h4(e15))) | ~ (h4(op1(e16, e14)) = op2(h4(e16), h4(e14))) | ~ (h4(op1(e16, e13)) = op2(h4(e16), h4(e13))) | ~ (h4(op1(e16, e12)) = op2(h4(e16), h4(e12))) | ~ (h4(op1(e16, e11)) = op2(h4(e16), h4(e11))) | ~ (h4(op1(e16, e10)) = op2(h4(e16), h4(e10))) | ~ (h4(op1(e15, e16)) = op2(h4(e15), h4(e16))) | ~ (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) | ~ (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) | ~ (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) | ~ (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) | ~ (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) | ~ (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) | ~ (h4(op1(e14, e16)) = op2(h4(e14), h4(e16))) | ~ (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e16)) = op2(h4(e13), h4(e16))) | ~ (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e16)) = op2(h4(e12), h4(e16))) | ~ (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e16)) = op2(h4(e11), h4(e16))) | ~ (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e16)) = op2(h4(e10), h4(e16))) | ~ (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e26 = h3(e16)) & ~ (e26 = h3(e15)) & ~ (e26 = h3(e14)) & ~ (e26 = h3(e13)) & ~ (e26 = h3(e12)) & ~ (e26 = h3(e11)) & ~ (e26 = h3(e10))) | (~ (e25 = h3(e16)) & ~ (e25 = h3(e15)) & ~ (e25 = h3(e14)) & ~ (e25 = h3(e13)) & ~ (e25 = h3(e12)) & ~ (e25 = h3(e11)) & ~ (e25 = h3(e10))) | (~ (e24 = h3(e16)) & ~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | (~ (e23 = h3(e16)) & ~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e16)) & ~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e16)) & ~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e16)) & ~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e16, e16)) = op2(h3(e16), h3(e16))) | ~ (h3(op1(e16, e15)) = op2(h3(e16), h3(e15))) | ~ (h3(op1(e16, e14)) = op2(h3(e16), h3(e14))) | ~ (h3(op1(e16, e13)) = op2(h3(e16), h3(e13))) | ~ (h3(op1(e16, e12)) = op2(h3(e16), h3(e12))) | ~ (h3(op1(e16, e11)) = op2(h3(e16), h3(e11))) | ~ (h3(op1(e16, e10)) = op2(h3(e16), h3(e10))) | ~ (h3(op1(e15, e16)) = op2(h3(e15), h3(e16))) | ~ (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) | ~ (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) | ~ (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) | ~ (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) | ~ (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) | ~ (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) | ~ (h3(op1(e14, e16)) = op2(h3(e14), h3(e16))) | ~ (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e16)) = op2(h3(e13), h3(e16))) | ~ (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e16)) = op2(h3(e12), h3(e16))) | ~ (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e16)) = op2(h3(e11), h3(e16))) | ~ (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e16)) = op2(h3(e10), h3(e16))) | ~ (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e26 = h2(e16)) & ~ (e26 = h2(e15)) & ~ (e26 = h2(e14)) & ~ (e26 = h2(e13)) & ~ (e26 = h2(e12)) & ~ (e26 = h2(e11)) & ~ (e26 = h2(e10))) | (~ (e25 = h2(e16)) & ~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10))) | (~ (e24 = h2(e16)) & ~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | (~ (e23 = h2(e16)) & ~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e16)) & ~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e16)) & ~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e16)) & ~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e16, e16)) = op2(h2(e16), h2(e16))) | ~ (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) | ~ (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e26 = h1(e16)) & ~ (e26 = h1(e15)) & ~ (e26 = h1(e14)) & ~ (e26 = h1(e13)) & ~ (e26 = h1(e12)) & ~ (e26 = h1(e11)) & ~ (e26 = h1(e10))) | (~ (e25 = h1(e16)) & ~ (e25 = h1(e15)) & ~ (e25 = h1(e14)) & ~ (e25 = h1(e13)) & ~ (e25 = h1(e12)) & ~ (e25 = h1(e11)) & ~ (e25 = h1(e10))) | (~ (e24 = h1(e16)) & ~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | (~ (e23 = h1(e16)) & ~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e16)) & ~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e16)) & ~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e16)) & ~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e16, e16)) = op2(h1(e16), h1(e16))) | ~ (h1(op1(e16, e15)) = op2(h1(e16), h1(e15))) | ~ (h1(op1(e16, e14)) = op2(h1(e16), h1(e14))) | ~ (h1(op1(e16, e13)) = op2(h1(e16), h1(e13))) | ~ (h1(op1(e16, e12)) = op2(h1(e16), h1(e12))) | ~ (h1(op1(e16, e11)) = op2(h1(e16), h1(e11))) | ~ (h1(op1(e16, e10)) = op2(h1(e16), h1(e10))) | ~ (h1(op1(e15, e16)) = op2(h1(e15), h1(e16))) | ~ (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) | ~ (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) | ~ (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) | ~ (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) | ~ (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) | ~ (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) | ~ (h1(op1(e14, e16)) = op2(h1(e14), h1(e16))) | ~ (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e16)) = op2(h1(e13), h1(e16))) | ~ (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e16)) = op2(h1(e12), h1(e16))) | ~ (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e16)) = op2(h1(e11), h1(e16))) | ~ (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e16)) = op2(h1(e10), h1(e16))) | ~ (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((((e26 = h7(e16)) | (e26 = h7(e15)) | (e26 = h7(e14)) | (e26 = h7(e13)) | (e26 = h7(e12)) | (e26 = h7(e11)) | (e26 = h7(e10))) & ((e25 = h7(e16)) | (e25 = h7(e15)) | (e25 = h7(e14)) | (e25 = h7(e13)) | (e25 = h7(e12)) | (e25 = h7(e11)) | (e25 = h7(e10))) & ((e24 = h7(e16)) | (e24 = h7(e15)) | (e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e16)) | (e23 = h7(e15)) | (e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e16)) | (e22 = h7(e15)) | (e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e16)) | (e21 = h7(e15)) | (e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e16)) | (e20 = h7(e15)) | (e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e16, e16)) = op2(h7(e16), h7(e16))) & (h7(op1(e16, e15)) = op2(h7(e16), h7(e15))) & (h7(op1(e16, e14)) = op2(h7(e16), h7(e14))) & (h7(op1(e16, e13)) = op2(h7(e16), h7(e13))) & (h7(op1(e16, e12)) = op2(h7(e16), h7(e12))) & (h7(op1(e16, e11)) = op2(h7(e16), h7(e11))) & (h7(op1(e16, e10)) = op2(h7(e16), h7(e10))) & (h7(op1(e15, e16)) = op2(h7(e15), h7(e16))) & (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) & (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) & (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) & (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) & (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) & (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) & (h7(op1(e14, e16)) = op2(h7(e14), h7(e16))) & (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e16)) = op2(h7(e13), h7(e16))) & (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e16)) = op2(h7(e12), h7(e16))) & (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e16)) = op2(h7(e11), h7(e16))) & (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e16)) = op2(h7(e10), h7(e16))) & (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e26 = h6(e16)) | (e26 = h6(e15)) | (e26 = h6(e14)) | (e26 = h6(e13)) | (e26 = h6(e12)) | (e26 = h6(e11)) | (e26 = h6(e10))) & ((e25 = h6(e16)) | (e25 = h6(e15)) | (e25 = h6(e14)) | (e25 = h6(e13)) | (e25 = h6(e12)) | (e25 = h6(e11)) | (e25 = h6(e10))) & ((e24 = h6(e16)) | (e24 = h6(e15)) | (e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e16)) | (e23 = h6(e15)) | (e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e16)) | (e22 = h6(e15)) | (e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e16)) | (e21 = h6(e15)) | (e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e16)) | (e20 = h6(e15)) | (e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) & (h6(op1(e16, e15)) = op2(h6(e16), h6(e15))) & (h6(op1(e16, e14)) = op2(h6(e16), h6(e14))) & (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) & (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) & (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) & (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) & (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) & (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) & (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) & (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) & (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) & (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) & (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) & (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) & (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) & (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) & (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) & (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) & (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e26 = h5(e16)) | (e26 = h5(e15)) | (e26 = h5(e14)) | (e26 = h5(e13)) | (e26 = h5(e12)) | (e26 = h5(e11)) | (e26 = h5(e10))) & ((e25 = h5(e16)) | (e25 = h5(e15)) | (e25 = h5(e14)) | (e25 = h5(e13)) | (e25 = h5(e12)) | (e25 = h5(e11)) | (e25 = h5(e10))) & ((e24 = h5(e16)) | (e24 = h5(e15)) | (e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e16)) | (e23 = h5(e15)) | (e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e16)) | (e22 = h5(e15)) | (e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e16)) | (e21 = h5(e15)) | (e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e16)) | (e20 = h5(e15)) | (e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e16, e16)) = op2(h5(e16), h5(e16))) & (h5(op1(e16, e15)) = op2(h5(e16), h5(e15))) & (h5(op1(e16, e14)) = op2(h5(e16), h5(e14))) & (h5(op1(e16, e13)) = op2(h5(e16), h5(e13))) & (h5(op1(e16, e12)) = op2(h5(e16), h5(e12))) & (h5(op1(e16, e11)) = op2(h5(e16), h5(e11))) & (h5(op1(e16, e10)) = op2(h5(e16), h5(e10))) & (h5(op1(e15, e16)) = op2(h5(e15), h5(e16))) & (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) & (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) & (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) & (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) & (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) & (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) & (h5(op1(e14, e16)) = op2(h5(e14), h5(e16))) & (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e16)) = op2(h5(e13), h5(e16))) & (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e16)) = op2(h5(e12), h5(e16))) & (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e16)) = op2(h5(e11), h5(e16))) & (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e16)) = op2(h5(e10), h5(e16))) & (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e26 = h4(e16)) | (e26 = h4(e15)) | (e26 = h4(e14)) | (e26 = h4(e13)) | (e26 = h4(e12)) | (e26 = h4(e11)) | (e26 = h4(e10))) & ((e25 = h4(e16)) | (e25 = h4(e15)) | (e25 = h4(e14)) | (e25 = h4(e13)) | (e25 = h4(e12)) | (e25 = h4(e11)) | (e25 = h4(e10))) & ((e24 = h4(e16)) | (e24 = h4(e15)) | (e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e16)) | (e23 = h4(e15)) | (e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e16)) | (e22 = h4(e15)) | (e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e16)) | (e21 = h4(e15)) | (e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e16)) | (e20 = h4(e15)) | (e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e16, e16)) = op2(h4(e16), h4(e16))) & (h4(op1(e16, e15)) = op2(h4(e16), h4(e15))) & (h4(op1(e16, e14)) = op2(h4(e16), h4(e14))) & (h4(op1(e16, e13)) = op2(h4(e16), h4(e13))) & (h4(op1(e16, e12)) = op2(h4(e16), h4(e12))) & (h4(op1(e16, e11)) = op2(h4(e16), h4(e11))) & (h4(op1(e16, e10)) = op2(h4(e16), h4(e10))) & (h4(op1(e15, e16)) = op2(h4(e15), h4(e16))) & (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) & (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) & (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) & (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) & (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) & (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) & (h4(op1(e14, e16)) = op2(h4(e14), h4(e16))) & (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e16)) = op2(h4(e13), h4(e16))) & (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e16)) = op2(h4(e12), h4(e16))) & (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e16)) = op2(h4(e11), h4(e16))) & (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e16)) = op2(h4(e10), h4(e16))) & (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e26 = h3(e16)) | (e26 = h3(e15)) | (e26 = h3(e14)) | (e26 = h3(e13)) | (e26 = h3(e12)) | (e26 = h3(e11)) | (e26 = h3(e10))) & ((e25 = h3(e16)) | (e25 = h3(e15)) | (e25 = h3(e14)) | (e25 = h3(e13)) | (e25 = h3(e12)) | (e25 = h3(e11)) | (e25 = h3(e10))) & ((e24 = h3(e16)) | (e24 = h3(e15)) | (e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e16)) | (e23 = h3(e15)) | (e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e16)) | (e22 = h3(e15)) | (e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e16)) | (e21 = h3(e15)) | (e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e16)) | (e20 = h3(e15)) | (e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e16, e16)) = op2(h3(e16), h3(e16))) & (h3(op1(e16, e15)) = op2(h3(e16), h3(e15))) & (h3(op1(e16, e14)) = op2(h3(e16), h3(e14))) & (h3(op1(e16, e13)) = op2(h3(e16), h3(e13))) & (h3(op1(e16, e12)) = op2(h3(e16), h3(e12))) & (h3(op1(e16, e11)) = op2(h3(e16), h3(e11))) & (h3(op1(e16, e10)) = op2(h3(e16), h3(e10))) & (h3(op1(e15, e16)) = op2(h3(e15), h3(e16))) & (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) & (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) & (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) & (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) & (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) & (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) & (h3(op1(e14, e16)) = op2(h3(e14), h3(e16))) & (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e16)) = op2(h3(e13), h3(e16))) & (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e16)) = op2(h3(e12), h3(e16))) & (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e16)) = op2(h3(e11), h3(e16))) & (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e16)) = op2(h3(e10), h3(e16))) & (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e26 = h2(e16)) | (e26 = h2(e15)) | (e26 = h2(e14)) | (e26 = h2(e13)) | (e26 = h2(e12)) | (e26 = h2(e11)) | (e26 = h2(e10))) & ((e25 = h2(e16)) | (e25 = h2(e15)) | (e25 = h2(e14)) | (e25 = h2(e13)) | (e25 = h2(e12)) | (e25 = h2(e11)) | (e25 = h2(e10))) & ((e24 = h2(e16)) | (e24 = h2(e15)) | (e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e16)) | (e23 = h2(e15)) | (e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e16)) | (e22 = h2(e15)) | (e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e16)) | (e21 = h2(e15)) | (e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e16)) | (e20 = h2(e15)) | (e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e16, e16)) = op2(h2(e16), h2(e16))) & (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) & (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) & (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) & (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) & (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) & (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) & (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) & (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) & (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) & (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) & (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) & (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) & (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) & (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) & (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) & (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) & (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) & (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) & (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e26 = h1(e16)) | (e26 = h1(e15)) | (e26 = h1(e14)) | (e26 = h1(e13)) | (e26 = h1(e12)) | (e26 = h1(e11)) | (e26 = h1(e10))) & ((e25 = h1(e16)) | (e25 = h1(e15)) | (e25 = h1(e14)) | (e25 = h1(e13)) | (e25 = h1(e12)) | (e25 = h1(e11)) | (e25 = h1(e10))) & ((e24 = h1(e16)) | (e24 = h1(e15)) | (e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e16)) | (e23 = h1(e15)) | (e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e16)) | (e22 = h1(e15)) | (e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e16)) | (e21 = h1(e15)) | (e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e16)) | (e20 = h1(e15)) | (e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e16, e16)) = op2(h1(e16), h1(e16))) & (h1(op1(e16, e15)) = op2(h1(e16), h1(e15))) & (h1(op1(e16, e14)) = op2(h1(e16), h1(e14))) & (h1(op1(e16, e13)) = op2(h1(e16), h1(e13))) & (h1(op1(e16, e12)) = op2(h1(e16), h1(e12))) & (h1(op1(e16, e11)) = op2(h1(e16), h1(e11))) & (h1(op1(e16, e10)) = op2(h1(e16), h1(e10))) & (h1(op1(e15, e16)) = op2(h1(e15), h1(e16))) & (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) & (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) & (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) & (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) & (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) & (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) & (h1(op1(e14, e16)) = op2(h1(e14), h1(e16))) & (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e16)) = op2(h1(e13), h1(e16))) & (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e16)) = op2(h1(e12), h1(e16))) & (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e16)) = op2(h1(e11), h1(e16))) & (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e16)) = op2(h1(e10), h1(e16))) & (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f23])).
fof(f23, plain, ~ ((((e26 = h7(e16)) | (e26 = h7(e15)) | (e26 = h7(e14)) | (e26 = h7(e13)) | (e26 = h7(e12)) | (e26 = h7(e11)) | (e26 = h7(e10))) & ((e25 = h7(e16)) | (e25 = h7(e15)) | (e25 = h7(e14)) | (e25 = h7(e13)) | (e25 = h7(e12)) | (e25 = h7(e11)) | (e25 = h7(e10))) & ((e24 = h7(e16)) | (e24 = h7(e15)) | (e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e16)) | (e23 = h7(e15)) | (e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e16)) | (e22 = h7(e15)) | (e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e16)) | (e21 = h7(e15)) | (e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e16)) | (e20 = h7(e15)) | (e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e16, e16)) = op2(h7(e16), h7(e16))) & (h7(op1(e16, e15)) = op2(h7(e16), h7(e15))) & (h7(op1(e16, e14)) = op2(h7(e16), h7(e14))) & (h7(op1(e16, e13)) = op2(h7(e16), h7(e13))) & (h7(op1(e16, e12)) = op2(h7(e16), h7(e12))) & (h7(op1(e16, e11)) = op2(h7(e16), h7(e11))) & (h7(op1(e16, e10)) = op2(h7(e16), h7(e10))) & (h7(op1(e15, e16)) = op2(h7(e15), h7(e16))) & (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) & (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) & (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) & (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) & (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) & (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) & (h7(op1(e14, e16)) = op2(h7(e14), h7(e16))) & (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e16)) = op2(h7(e13), h7(e16))) & (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e16)) = op2(h7(e12), h7(e16))) & (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e16)) = op2(h7(e11), h7(e16))) & (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e16)) = op2(h7(e10), h7(e16))) & (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e26 = h6(e16)) | (e26 = h6(e15)) | (e26 = h6(e14)) | (e26 = h6(e13)) | (e26 = h6(e12)) | (e26 = h6(e11)) | (e26 = h6(e10))) & ((e25 = h6(e16)) | (e25 = h6(e15)) | (e25 = h6(e14)) | (e25 = h6(e13)) | (e25 = h6(e12)) | (e25 = h6(e11)) | (e25 = h6(e10))) & ((e24 = h6(e16)) | (e24 = h6(e15)) | (e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e16)) | (e23 = h6(e15)) | (e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e16)) | (e22 = h6(e15)) | (e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e16)) | (e21 = h6(e15)) | (e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e16)) | (e20 = h6(e15)) | (e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) & (h6(op1(e16, e15)) = op2(h6(e16), h6(e15))) & (h6(op1(e16, e14)) = op2(h6(e16), h6(e14))) & (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) & (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) & (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) & (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) & (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) & (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) & (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) & (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) & (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) & (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) & (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) & (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) & (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) & (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) & (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) & (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) & (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e26 = h5(e16)) | (e26 = h5(e15)) | (e26 = h5(e14)) | (e26 = h5(e13)) | (e26 = h5(e12)) | (e26 = h5(e11)) | (e26 = h5(e10))) & ((e25 = h5(e16)) | (e25 = h5(e15)) | (e25 = h5(e14)) | (e25 = h5(e13)) | (e25 = h5(e12)) | (e25 = h5(e11)) | (e25 = h5(e10))) & ((e24 = h5(e16)) | (e24 = h5(e15)) | (e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e16)) | (e23 = h5(e15)) | (e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e16)) | (e22 = h5(e15)) | (e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e16)) | (e21 = h5(e15)) | (e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e16)) | (e20 = h5(e15)) | (e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e16, e16)) = op2(h5(e16), h5(e16))) & (h5(op1(e16, e15)) = op2(h5(e16), h5(e15))) & (h5(op1(e16, e14)) = op2(h5(e16), h5(e14))) & (h5(op1(e16, e13)) = op2(h5(e16), h5(e13))) & (h5(op1(e16, e12)) = op2(h5(e16), h5(e12))) & (h5(op1(e16, e11)) = op2(h5(e16), h5(e11))) & (h5(op1(e16, e10)) = op2(h5(e16), h5(e10))) & (h5(op1(e15, e16)) = op2(h5(e15), h5(e16))) & (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) & (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) & (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) & (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) & (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) & (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) & (h5(op1(e14, e16)) = op2(h5(e14), h5(e16))) & (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e16)) = op2(h5(e13), h5(e16))) & (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e16)) = op2(h5(e12), h5(e16))) & (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e16)) = op2(h5(e11), h5(e16))) & (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e16)) = op2(h5(e10), h5(e16))) & (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e26 = h4(e16)) | (e26 = h4(e15)) | (e26 = h4(e14)) | (e26 = h4(e13)) | (e26 = h4(e12)) | (e26 = h4(e11)) | (e26 = h4(e10))) & ((e25 = h4(e16)) | (e25 = h4(e15)) | (e25 = h4(e14)) | (e25 = h4(e13)) | (e25 = h4(e12)) | (e25 = h4(e11)) | (e25 = h4(e10))) & ((e24 = h4(e16)) | (e24 = h4(e15)) | (e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e16)) | (e23 = h4(e15)) | (e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e16)) | (e22 = h4(e15)) | (e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e16)) | (e21 = h4(e15)) | (e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e16)) | (e20 = h4(e15)) | (e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e16, e16)) = op2(h4(e16), h4(e16))) & (h4(op1(e16, e15)) = op2(h4(e16), h4(e15))) & (h4(op1(e16, e14)) = op2(h4(e16), h4(e14))) & (h4(op1(e16, e13)) = op2(h4(e16), h4(e13))) & (h4(op1(e16, e12)) = op2(h4(e16), h4(e12))) & (h4(op1(e16, e11)) = op2(h4(e16), h4(e11))) & (h4(op1(e16, e10)) = op2(h4(e16), h4(e10))) & (h4(op1(e15, e16)) = op2(h4(e15), h4(e16))) & (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) & (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) & (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) & (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) & (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) & (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) & (h4(op1(e14, e16)) = op2(h4(e14), h4(e16))) & (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e16)) = op2(h4(e13), h4(e16))) & (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e16)) = op2(h4(e12), h4(e16))) & (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e16)) = op2(h4(e11), h4(e16))) & (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e16)) = op2(h4(e10), h4(e16))) & (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e26 = h3(e16)) | (e26 = h3(e15)) | (e26 = h3(e14)) | (e26 = h3(e13)) | (e26 = h3(e12)) | (e26 = h3(e11)) | (e26 = h3(e10))) & ((e25 = h3(e16)) | (e25 = h3(e15)) | (e25 = h3(e14)) | (e25 = h3(e13)) | (e25 = h3(e12)) | (e25 = h3(e11)) | (e25 = h3(e10))) & ((e24 = h3(e16)) | (e24 = h3(e15)) | (e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e16)) | (e23 = h3(e15)) | (e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e16)) | (e22 = h3(e15)) | (e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e16)) | (e21 = h3(e15)) | (e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e16)) | (e20 = h3(e15)) | (e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e16, e16)) = op2(h3(e16), h3(e16))) & (h3(op1(e16, e15)) = op2(h3(e16), h3(e15))) & (h3(op1(e16, e14)) = op2(h3(e16), h3(e14))) & (h3(op1(e16, e13)) = op2(h3(e16), h3(e13))) & (h3(op1(e16, e12)) = op2(h3(e16), h3(e12))) & (h3(op1(e16, e11)) = op2(h3(e16), h3(e11))) & (h3(op1(e16, e10)) = op2(h3(e16), h3(e10))) & (h3(op1(e15, e16)) = op2(h3(e15), h3(e16))) & (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) & (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) & (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) & (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) & (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) & (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) & (h3(op1(e14, e16)) = op2(h3(e14), h3(e16))) & (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e16)) = op2(h3(e13), h3(e16))) & (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e16)) = op2(h3(e12), h3(e16))) & (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e16)) = op2(h3(e11), h3(e16))) & (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e16)) = op2(h3(e10), h3(e16))) & (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e26 = h2(e16)) | (e26 = h2(e15)) | (e26 = h2(e14)) | (e26 = h2(e13)) | (e26 = h2(e12)) | (e26 = h2(e11)) | (e26 = h2(e10))) & ((e25 = h2(e16)) | (e25 = h2(e15)) | (e25 = h2(e14)) | (e25 = h2(e13)) | (e25 = h2(e12)) | (e25 = h2(e11)) | (e25 = h2(e10))) & ((e24 = h2(e16)) | (e24 = h2(e15)) | (e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e16)) | (e23 = h2(e15)) | (e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e16)) | (e22 = h2(e15)) | (e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e16)) | (e21 = h2(e15)) | (e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e16)) | (e20 = h2(e15)) | (e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e16, e16)) = op2(h2(e16), h2(e16))) & (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) & (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) & (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) & (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) & (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) & (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) & (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) & (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) & (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) & (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) & (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) & (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) & (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) & (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) & (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) & (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) & (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) & (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) & (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e26 = h1(e16)) | (e26 = h1(e15)) | (e26 = h1(e14)) | (e26 = h1(e13)) | (e26 = h1(e12)) | (e26 = h1(e11)) | (e26 = h1(e10))) & ((e25 = h1(e16)) | (e25 = h1(e15)) | (e25 = h1(e14)) | (e25 = h1(e13)) | (e25 = h1(e12)) | (e25 = h1(e11)) | (e25 = h1(e10))) & ((e24 = h1(e16)) | (e24 = h1(e15)) | (e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e16)) | (e23 = h1(e15)) | (e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e16)) | (e22 = h1(e15)) | (e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e16)) | (e21 = h1(e15)) | (e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e16)) | (e20 = h1(e15)) | (e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e16, e16)) = op2(h1(e16), h1(e16))) & (h1(op1(e16, e15)) = op2(h1(e16), h1(e15))) & (h1(op1(e16, e14)) = op2(h1(e16), h1(e14))) & (h1(op1(e16, e13)) = op2(h1(e16), h1(e13))) & (h1(op1(e16, e12)) = op2(h1(e16), h1(e12))) & (h1(op1(e16, e11)) = op2(h1(e16), h1(e11))) & (h1(op1(e16, e10)) = op2(h1(e16), h1(e10))) & (h1(op1(e15, e16)) = op2(h1(e15), h1(e16))) & (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) & (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) & (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) & (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) & (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) & (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) & (h1(op1(e14, e16)) = op2(h1(e14), h1(e16))) & (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e16)) = op2(h1(e13), h1(e16))) & (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e16)) = op2(h1(e12), h1(e16))) & (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e16)) = op2(h1(e11), h1(e16))) & (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e16)) = op2(h1(e10), h1(e16))) & (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', co1)).
fof(f5820, plain, (~ spl42_879 | ~ spl42_883), inference(avatar_split_clause, [], [f1494, f5817, f5798])).
fof(f1494, plain, (~ (e20 = h2(e12)) | ~ sP6), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e20 = h2(e16)) & ~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP6), inference(nnf_transformation, [], [f32])).
fof(f5779, plain, (~ spl42_872 | ~ spl42_875), inference(avatar_split_clause, [], [f1481, f5776, f5762])).
fof(f1481, plain, (~ (e22 = h2(e13)) | ~ sP8), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e22 = h2(e16)) & ~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP8), inference(nnf_transformation, [], [f34])).
fof(f5759, plain, (~ spl42_865 | ~ spl42_871), inference(avatar_split_clause, [], [f1471, f5756, f5727])).
fof(f1471, plain, (~ (e23 = h2(e10)) | ~ sP9), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e23 = h2(e16)) & ~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP9), inference(nnf_transformation, [], [f35])).
fof(f5704, plain, (~ spl42_858 | ~ spl42_860), inference(avatar_split_clause, [], [f1468, f5701, f5692])).
fof(f1468, plain, (~ (e24 = h2(e14)) | ~ sP10), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e24 = h2(e16)) & ~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | ~ sP10), inference(nnf_transformation, [], [f36])).
fof(f5684, plain, (~ spl42_851 | ~ spl42_856), inference(avatar_split_clause, [], [f1458, f5681, f5657])).
fof(f1458, plain, (~ (e25 = h2(e11)) | ~ sP11), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e25 = h2(e16)) & ~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10))) | ~ sP11), inference(nnf_transformation, [], [f37])).
fof(f4643, plain, spl42_382, inference(avatar_split_clause, [], [f4642, f3267])).
fof(f4642, plain, (e23 = op2(e26, e21)), inference(backward_demodulation, [], [f1195, f1193])).
fof(f1195, plain, (e23 = op2(e26, op2(op2(e26, e26), op2(e26, e26)))), inference(cnf_transformation, [], [f15])).
fof(f4640, plain, spl42_369, inference(avatar_split_clause, [], [f4639, f3213])).
fof(f4639, plain, (e24 = op2(e26, e23)), inference(backward_demodulation, [], [f1196, f1195])).
fof(f1196, plain, (e24 = op2(e26, op2(e26, op2(op2(e26, e26), op2(e26, e26))))), inference(cnf_transformation, [], [f15])).
fof(f4638, plain, spl42_349, inference(avatar_split_clause, [], [f1197, f3130])).
fof(f1197, plain, (e25 = op2(e26, e26)), inference(cnf_transformation, [], [f15])).
fof(f4637, plain, spl42_39, inference(avatar_split_clause, [], [f4636, f1748])).
fof(f4636, plain, (e13 = op1(e16, e11)), inference(backward_demodulation, [], [f1189, f1187])).
fof(f1189, plain, (e13 = op1(e16, op1(op1(e16, e16), op1(e16, e16)))), inference(cnf_transformation, [], [f14])).
fof(f4634, plain, spl42_26, inference(avatar_split_clause, [], [f4633, f1694])).
fof(f4633, plain, (e14 = op1(e16, e13)), inference(backward_demodulation, [], [f1190, f1189])).
fof(f1190, plain, (e14 = op1(e16, op1(e16, op1(op1(e16, e16), op1(e16, e16))))), inference(cnf_transformation, [], [f14])).
fof(f4632, plain, spl42_6, inference(avatar_split_clause, [], [f1191, f1611])).
fof(f1191, plain, (e15 = op1(e16, e16)), inference(cnf_transformation, [], [f14])).
fof(f4631, plain, (spl42_680 | spl42_625 | spl42_570 | spl42_515 | spl42_460 | spl42_405 | spl42_350), inference(avatar_split_clause, [], [f1184, f3134, f3362, f3590, f3818, f4046, f4274, f4502])).
fof(f1184, plain, ((e26 = op2(e26, e26)) | (e25 = op2(e25, e25)) | (e24 = op2(e24, e24)) | (e23 = op2(e23, e23)) | (e22 = op2(e22, e22)) | (e21 = op2(e21, e21)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((~ (e26 = op2(e26, e26)) | ~ (e25 = op2(e25, e25)) | ~ (e24 = op2(e24, e24)) | ~ (e23 = op2(e23, e23)) | ~ (e22 = op2(e22, e22)) | ~ (e21 = op2(e21, e21)) | ~ (e20 = op2(e20, e20))) & ((e26 = op2(e26, e26)) | (e25 = op2(e25, e25)) | (e24 = op2(e24, e24)) | (e23 = op2(e23, e23)) | (e22 = op2(e22, e22)) | (e21 = op2(e21, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax13)).
fof(f4629, plain, (spl42_337 | spl42_282 | spl42_227 | spl42_172 | spl42_117 | spl42_62 | spl42_7), inference(avatar_split_clause, [], [f1182, f1615, f1843, f2071, f2299, f2527, f2755, f2983])).
fof(f1182, plain, ((e16 = op1(e16, e16)) | (e15 = op1(e15, e15)) | (e14 = op1(e14, e14)) | (e13 = op1(e13, e13)) | (e12 = op1(e12, e12)) | (e11 = op1(e11, e11)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((~ (e16 = op1(e16, e16)) | ~ (e15 = op1(e15, e15)) | ~ (e14 = op1(e14, e14)) | ~ (e13 = op1(e13, e13)) | ~ (e12 = op1(e12, e12)) | ~ (e11 = op1(e11, e11)) | ~ (e10 = op1(e10, e10))) & ((e16 = op1(e16, e16)) | (e15 = op1(e15, e15)) | (e14 = op1(e14, e14)) | (e13 = op1(e13, e13)) | (e12 = op1(e12, e12)) | (e11 = op1(e11, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax12)).
fof(f4623, plain, (spl42_682 | spl42_675 | spl42_668 | spl42_661 | spl42_654 | spl42_647 | spl42_640), inference(avatar_split_clause, [], [f360, f4336, f4365, f4394, f4423, f4452, f4481, f4510])).
fof(f360, plain, ((e22 = op2(e20, e26)) | (e22 = op2(e20, e25)) | (e22 = op2(e20, e24)) | (e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e26 = op2(e26, e26)) | (e26 = op2(e25, e26)) | (e26 = op2(e24, e26)) | (e26 = op2(e23, e26)) | (e26 = op2(e22, e26)) | (e26 = op2(e21, e26)) | (e26 = op2(e20, e26))) & ((e26 = op2(e26, e26)) | (e26 = op2(e26, e25)) | (e26 = op2(e26, e24)) | (e26 = op2(e26, e23)) | (e26 = op2(e26, e22)) | (e26 = op2(e26, e21)) | (e26 = op2(e26, e20))) & ((e25 = op2(e26, e26)) | (e25 = op2(e25, e26)) | (e25 = op2(e24, e26)) | (e25 = op2(e23, e26)) | (e25 = op2(e22, e26)) | (e25 = op2(e21, e26)) | (e25 = op2(e20, e26))) & ((e25 = op2(e26, e26)) | (e25 = op2(e26, e25)) | (e25 = op2(e26, e24)) | (e25 = op2(e26, e23)) | (e25 = op2(e26, e22)) | (e25 = op2(e26, e21)) | (e25 = op2(e26, e20))) & ((e24 = op2(e26, e26)) | (e24 = op2(e25, e26)) | (e24 = op2(e24, e26)) | (e24 = op2(e23, e26)) | (e24 = op2(e22, e26)) | (e24 = op2(e21, e26)) | (e24 = op2(e20, e26))) & ((e24 = op2(e26, e26)) | (e24 = op2(e26, e25)) | (e24 = op2(e26, e24)) | (e24 = op2(e26, e23)) | (e24 = op2(e26, e22)) | (e24 = op2(e26, e21)) | (e24 = op2(e26, e20))) & ((e23 = op2(e26, e26)) | (e23 = op2(e25, e26)) | (e23 = op2(e24, e26)) | (e23 = op2(e23, e26)) | (e23 = op2(e22, e26)) | (e23 = op2(e21, e26)) | (e23 = op2(e20, e26))) & ((e23 = op2(e26, e26)) | (e23 = op2(e26, e25)) | (e23 = op2(e26, e24)) | (e23 = op2(e26, e23)) | (e23 = op2(e26, e22)) | (e23 = op2(e26, e21)) | (e23 = op2(e26, e20))) & ((e22 = op2(e26, e26)) | (e22 = op2(e25, e26)) | (e22 = op2(e24, e26)) | (e22 = op2(e23, e26)) | (e22 = op2(e22, e26)) | (e22 = op2(e21, e26)) | (e22 = op2(e20, e26))) & ((e22 = op2(e26, e26)) | (e22 = op2(e26, e25)) | (e22 = op2(e26, e24)) | (e22 = op2(e26, e23)) | (e22 = op2(e26, e22)) | (e22 = op2(e26, e21)) | (e22 = op2(e26, e20))) & ((e21 = op2(e26, e26)) | (e21 = op2(e25, e26)) | (e21 = op2(e24, e26)) | (e21 = op2(e23, e26)) | (e21 = op2(e22, e26)) | (e21 = op2(e21, e26)) | (e21 = op2(e20, e26))) & ((e21 = op2(e26, e26)) | (e21 = op2(e26, e25)) | (e21 = op2(e26, e24)) | (e21 = op2(e26, e23)) | (e21 = op2(e26, e22)) | (e21 = op2(e26, e21)) | (e21 = op2(e26, e20))) & ((e20 = op2(e26, e26)) | (e20 = op2(e25, e26)) | (e20 = op2(e24, e26)) | (e20 = op2(e23, e26)) | (e20 = op2(e22, e26)) | (e20 = op2(e21, e26)) | (e20 = op2(e20, e26))) & ((e20 = op2(e26, e26)) | (e20 = op2(e26, e25)) | (e20 = op2(e26, e24)) | (e20 = op2(e26, e23)) | (e20 = op2(e26, e22)) | (e20 = op2(e26, e21)) | (e20 = op2(e26, e20))) & ((e26 = op2(e26, e25)) | (e26 = op2(e25, e25)) | (e26 = op2(e24, e25)) | (e26 = op2(e23, e25)) | (e26 = op2(e22, e25)) | (e26 = op2(e21, e25)) | (e26 = op2(e20, e25))) & ((e26 = op2(e25, e26)) | (e26 = op2(e25, e25)) | (e26 = op2(e25, e24)) | (e26 = op2(e25, e23)) | (e26 = op2(e25, e22)) | (e26 = op2(e25, e21)) | (e26 = op2(e25, e20))) & ((e25 = op2(e26, e25)) | (e25 = op2(e25, e25)) | (e25 = op2(e24, e25)) | (e25 = op2(e23, e25)) | (e25 = op2(e22, e25)) | (e25 = op2(e21, e25)) | (e25 = op2(e20, e25))) & ((e25 = op2(e25, e26)) | (e25 = op2(e25, e25)) | (e25 = op2(e25, e24)) | (e25 = op2(e25, e23)) | (e25 = op2(e25, e22)) | (e25 = op2(e25, e21)) | (e25 = op2(e25, e20))) & ((e24 = op2(e26, e25)) | (e24 = op2(e25, e25)) | (e24 = op2(e24, e25)) | (e24 = op2(e23, e25)) | (e24 = op2(e22, e25)) | (e24 = op2(e21, e25)) | (e24 = op2(e20, e25))) & ((e24 = op2(e25, e26)) | (e24 = op2(e25, e25)) | (e24 = op2(e25, e24)) | (e24 = op2(e25, e23)) | (e24 = op2(e25, e22)) | (e24 = op2(e25, e21)) | (e24 = op2(e25, e20))) & ((e23 = op2(e26, e25)) | (e23 = op2(e25, e25)) | (e23 = op2(e24, e25)) | (e23 = op2(e23, e25)) | (e23 = op2(e22, e25)) | (e23 = op2(e21, e25)) | (e23 = op2(e20, e25))) & ((e23 = op2(e25, e26)) | (e23 = op2(e25, e25)) | (e23 = op2(e25, e24)) | (e23 = op2(e25, e23)) | (e23 = op2(e25, e22)) | (e23 = op2(e25, e21)) | (e23 = op2(e25, e20))) & ((e22 = op2(e26, e25)) | (e22 = op2(e25, e25)) | (e22 = op2(e24, e25)) | (e22 = op2(e23, e25)) | (e22 = op2(e22, e25)) | (e22 = op2(e21, e25)) | (e22 = op2(e20, e25))) & ((e22 = op2(e25, e26)) | (e22 = op2(e25, e25)) | (e22 = op2(e25, e24)) | (e22 = op2(e25, e23)) | (e22 = op2(e25, e22)) | (e22 = op2(e25, e21)) | (e22 = op2(e25, e20))) & ((e21 = op2(e26, e25)) | (e21 = op2(e25, e25)) | (e21 = op2(e24, e25)) | (e21 = op2(e23, e25)) | (e21 = op2(e22, e25)) | (e21 = op2(e21, e25)) | (e21 = op2(e20, e25))) & ((e21 = op2(e25, e26)) | (e21 = op2(e25, e25)) | (e21 = op2(e25, e24)) | (e21 = op2(e25, e23)) | (e21 = op2(e25, e22)) | (e21 = op2(e25, e21)) | (e21 = op2(e25, e20))) & ((e20 = op2(e26, e25)) | (e20 = op2(e25, e25)) | (e20 = op2(e24, e25)) | (e20 = op2(e23, e25)) | (e20 = op2(e22, e25)) | (e20 = op2(e21, e25)) | (e20 = op2(e20, e25))) & ((e20 = op2(e25, e26)) | (e20 = op2(e25, e25)) | (e20 = op2(e25, e24)) | (e20 = op2(e25, e23)) | (e20 = op2(e25, e22)) | (e20 = op2(e25, e21)) | (e20 = op2(e25, e20))) & ((e26 = op2(e26, e24)) | (e26 = op2(e25, e24)) | (e26 = op2(e24, e24)) | (e26 = op2(e23, e24)) | (e26 = op2(e22, e24)) | (e26 = op2(e21, e24)) | (e26 = op2(e20, e24))) & ((e26 = op2(e24, e26)) | (e26 = op2(e24, e25)) | (e26 = op2(e24, e24)) | (e26 = op2(e24, e23)) | (e26 = op2(e24, e22)) | (e26 = op2(e24, e21)) | (e26 = op2(e24, e20))) & ((e25 = op2(e26, e24)) | (e25 = op2(e25, e24)) | (e25 = op2(e24, e24)) | (e25 = op2(e23, e24)) | (e25 = op2(e22, e24)) | (e25 = op2(e21, e24)) | (e25 = op2(e20, e24))) & ((e25 = op2(e24, e26)) | (e25 = op2(e24, e25)) | (e25 = op2(e24, e24)) | (e25 = op2(e24, e23)) | (e25 = op2(e24, e22)) | (e25 = op2(e24, e21)) | (e25 = op2(e24, e20))) & ((e24 = op2(e26, e24)) | (e24 = op2(e25, e24)) | (e24 = op2(e24, e24)) | (e24 = op2(e23, e24)) | (e24 = op2(e22, e24)) | (e24 = op2(e21, e24)) | (e24 = op2(e20, e24))) & ((e24 = op2(e24, e26)) | (e24 = op2(e24, e25)) | (e24 = op2(e24, e24)) | (e24 = op2(e24, e23)) | (e24 = op2(e24, e22)) | (e24 = op2(e24, e21)) | (e24 = op2(e24, e20))) & ((e23 = op2(e26, e24)) | (e23 = op2(e25, e24)) | (e23 = op2(e24, e24)) | (e23 = op2(e23, e24)) | (e23 = op2(e22, e24)) | (e23 = op2(e21, e24)) | (e23 = op2(e20, e24))) & ((e23 = op2(e24, e26)) | (e23 = op2(e24, e25)) | (e23 = op2(e24, e24)) | (e23 = op2(e24, e23)) | (e23 = op2(e24, e22)) | (e23 = op2(e24, e21)) | (e23 = op2(e24, e20))) & ((e22 = op2(e26, e24)) | (e22 = op2(e25, e24)) | (e22 = op2(e24, e24)) | (e22 = op2(e23, e24)) | (e22 = op2(e22, e24)) | (e22 = op2(e21, e24)) | (e22 = op2(e20, e24))) & ((e22 = op2(e24, e26)) | (e22 = op2(e24, e25)) | (e22 = op2(e24, e24)) | (e22 = op2(e24, e23)) | (e22 = op2(e24, e22)) | (e22 = op2(e24, e21)) | (e22 = op2(e24, e20))) & ((e21 = op2(e26, e24)) | (e21 = op2(e25, e24)) | (e21 = op2(e24, e24)) | (e21 = op2(e23, e24)) | (e21 = op2(e22, e24)) | (e21 = op2(e21, e24)) | (e21 = op2(e20, e24))) & ((e21 = op2(e24, e26)) | (e21 = op2(e24, e25)) | (e21 = op2(e24, e24)) | (e21 = op2(e24, e23)) | (e21 = op2(e24, e22)) | (e21 = op2(e24, e21)) | (e21 = op2(e24, e20))) & ((e20 = op2(e26, e24)) | (e20 = op2(e25, e24)) | (e20 = op2(e24, e24)) | (e20 = op2(e23, e24)) | (e20 = op2(e22, e24)) | (e20 = op2(e21, e24)) | (e20 = op2(e20, e24))) & ((e20 = op2(e24, e26)) | (e20 = op2(e24, e25)) | (e20 = op2(e24, e24)) | (e20 = op2(e24, e23)) | (e20 = op2(e24, e22)) | (e20 = op2(e24, e21)) | (e20 = op2(e24, e20))) & ((e26 = op2(e26, e23)) | (e26 = op2(e25, e23)) | (e26 = op2(e24, e23)) | (e26 = op2(e23, e23)) | (e26 = op2(e22, e23)) | (e26 = op2(e21, e23)) | (e26 = op2(e20, e23))) & ((e26 = op2(e23, e26)) | (e26 = op2(e23, e25)) | (e26 = op2(e23, e24)) | (e26 = op2(e23, e23)) | (e26 = op2(e23, e22)) | (e26 = op2(e23, e21)) | (e26 = op2(e23, e20))) & ((e25 = op2(e26, e23)) | (e25 = op2(e25, e23)) | (e25 = op2(e24, e23)) | (e25 = op2(e23, e23)) | (e25 = op2(e22, e23)) | (e25 = op2(e21, e23)) | (e25 = op2(e20, e23))) & ((e25 = op2(e23, e26)) | (e25 = op2(e23, e25)) | (e25 = op2(e23, e24)) | (e25 = op2(e23, e23)) | (e25 = op2(e23, e22)) | (e25 = op2(e23, e21)) | (e25 = op2(e23, e20))) & ((e24 = op2(e26, e23)) | (e24 = op2(e25, e23)) | (e24 = op2(e24, e23)) | (e24 = op2(e23, e23)) | (e24 = op2(e22, e23)) | (e24 = op2(e21, e23)) | (e24 = op2(e20, e23))) & ((e24 = op2(e23, e26)) | (e24 = op2(e23, e25)) | (e24 = op2(e23, e24)) | (e24 = op2(e23, e23)) | (e24 = op2(e23, e22)) | (e24 = op2(e23, e21)) | (e24 = op2(e23, e20))) & ((e23 = op2(e26, e23)) | (e23 = op2(e25, e23)) | (e23 = op2(e24, e23)) | (e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e26)) | (e23 = op2(e23, e25)) | (e23 = op2(e23, e24)) | (e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e26, e23)) | (e22 = op2(e25, e23)) | (e22 = op2(e24, e23)) | (e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e26)) | (e22 = op2(e23, e25)) | (e22 = op2(e23, e24)) | (e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e26, e23)) | (e21 = op2(e25, e23)) | (e21 = op2(e24, e23)) | (e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e26)) | (e21 = op2(e23, e25)) | (e21 = op2(e23, e24)) | (e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e26, e23)) | (e20 = op2(e25, e23)) | (e20 = op2(e24, e23)) | (e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e26)) | (e20 = op2(e23, e25)) | (e20 = op2(e23, e24)) | (e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e26 = op2(e26, e22)) | (e26 = op2(e25, e22)) | (e26 = op2(e24, e22)) | (e26 = op2(e23, e22)) | (e26 = op2(e22, e22)) | (e26 = op2(e21, e22)) | (e26 = op2(e20, e22))) & ((e26 = op2(e22, e26)) | (e26 = op2(e22, e25)) | (e26 = op2(e22, e24)) | (e26 = op2(e22, e23)) | (e26 = op2(e22, e22)) | (e26 = op2(e22, e21)) | (e26 = op2(e22, e20))) & ((e25 = op2(e26, e22)) | (e25 = op2(e25, e22)) | (e25 = op2(e24, e22)) | (e25 = op2(e23, e22)) | (e25 = op2(e22, e22)) | (e25 = op2(e21, e22)) | (e25 = op2(e20, e22))) & ((e25 = op2(e22, e26)) | (e25 = op2(e22, e25)) | (e25 = op2(e22, e24)) | (e25 = op2(e22, e23)) | (e25 = op2(e22, e22)) | (e25 = op2(e22, e21)) | (e25 = op2(e22, e20))) & ((e24 = op2(e26, e22)) | (e24 = op2(e25, e22)) | (e24 = op2(e24, e22)) | (e24 = op2(e23, e22)) | (e24 = op2(e22, e22)) | (e24 = op2(e21, e22)) | (e24 = op2(e20, e22))) & ((e24 = op2(e22, e26)) | (e24 = op2(e22, e25)) | (e24 = op2(e22, e24)) | (e24 = op2(e22, e23)) | (e24 = op2(e22, e22)) | (e24 = op2(e22, e21)) | (e24 = op2(e22, e20))) & ((e23 = op2(e26, e22)) | (e23 = op2(e25, e22)) | (e23 = op2(e24, e22)) | (e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e26)) | (e23 = op2(e22, e25)) | (e23 = op2(e22, e24)) | (e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e26, e22)) | (e22 = op2(e25, e22)) | (e22 = op2(e24, e22)) | (e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e26)) | (e22 = op2(e22, e25)) | (e22 = op2(e22, e24)) | (e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e26, e22)) | (e21 = op2(e25, e22)) | (e21 = op2(e24, e22)) | (e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e26)) | (e21 = op2(e22, e25)) | (e21 = op2(e22, e24)) | (e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e26, e22)) | (e20 = op2(e25, e22)) | (e20 = op2(e24, e22)) | (e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e26)) | (e20 = op2(e22, e25)) | (e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e26 = op2(e26, e21)) | (e26 = op2(e25, e21)) | (e26 = op2(e24, e21)) | (e26 = op2(e23, e21)) | (e26 = op2(e22, e21)) | (e26 = op2(e21, e21)) | (e26 = op2(e20, e21))) & ((e26 = op2(e21, e26)) | (e26 = op2(e21, e25)) | (e26 = op2(e21, e24)) | (e26 = op2(e21, e23)) | (e26 = op2(e21, e22)) | (e26 = op2(e21, e21)) | (e26 = op2(e21, e20))) & ((e25 = op2(e26, e21)) | (e25 = op2(e25, e21)) | (e25 = op2(e24, e21)) | (e25 = op2(e23, e21)) | (e25 = op2(e22, e21)) | (e25 = op2(e21, e21)) | (e25 = op2(e20, e21))) & ((e25 = op2(e21, e26)) | (e25 = op2(e21, e25)) | (e25 = op2(e21, e24)) | (e25 = op2(e21, e23)) | (e25 = op2(e21, e22)) | (e25 = op2(e21, e21)) | (e25 = op2(e21, e20))) & ((e24 = op2(e26, e21)) | (e24 = op2(e25, e21)) | (e24 = op2(e24, e21)) | (e24 = op2(e23, e21)) | (e24 = op2(e22, e21)) | (e24 = op2(e21, e21)) | (e24 = op2(e20, e21))) & ((e24 = op2(e21, e26)) | (e24 = op2(e21, e25)) | (e24 = op2(e21, e24)) | (e24 = op2(e21, e23)) | (e24 = op2(e21, e22)) | (e24 = op2(e21, e21)) | (e24 = op2(e21, e20))) & ((e23 = op2(e26, e21)) | (e23 = op2(e25, e21)) | (e23 = op2(e24, e21)) | (e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e26)) | (e23 = op2(e21, e25)) | (e23 = op2(e21, e24)) | (e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e26, e21)) | (e22 = op2(e25, e21)) | (e22 = op2(e24, e21)) | (e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e26)) | (e22 = op2(e21, e25)) | (e22 = op2(e21, e24)) | (e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e26, e21)) | (e21 = op2(e25, e21)) | (e21 = op2(e24, e21)) | (e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e26)) | (e21 = op2(e21, e25)) | (e21 = op2(e21, e24)) | (e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e26, e21)) | (e20 = op2(e25, e21)) | (e20 = op2(e24, e21)) | (e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e26)) | (e20 = op2(e21, e25)) | (e20 = op2(e21, e24)) | (e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e26 = op2(e26, e20)) | (e26 = op2(e25, e20)) | (e26 = op2(e24, e20)) | (e26 = op2(e23, e20)) | (e26 = op2(e22, e20)) | (e26 = op2(e21, e20)) | (op2(e20, e20) = e26)) & ((e26 = op2(e20, e26)) | (e26 = op2(e20, e25)) | (e26 = op2(e20, e24)) | (e26 = op2(e20, e23)) | (e26 = op2(e20, e22)) | (e26 = op2(e20, e21)) | (op2(e20, e20) = e26)) & ((e25 = op2(e26, e20)) | (e25 = op2(e25, e20)) | (e25 = op2(e24, e20)) | (e25 = op2(e23, e20)) | (e25 = op2(e22, e20)) | (e25 = op2(e21, e20)) | (op2(e20, e20) = e25)) & ((e25 = op2(e20, e26)) | (e25 = op2(e20, e25)) | (e25 = op2(e20, e24)) | (e25 = op2(e20, e23)) | (e25 = op2(e20, e22)) | (e25 = op2(e20, e21)) | (op2(e20, e20) = e25)) & ((e24 = op2(e26, e20)) | (e24 = op2(e25, e20)) | (e24 = op2(e24, e20)) | (e24 = op2(e23, e20)) | (e24 = op2(e22, e20)) | (e24 = op2(e21, e20)) | (op2(e20, e20) = e24)) & ((e24 = op2(e20, e26)) | (e24 = op2(e20, e25)) | (e24 = op2(e20, e24)) | (e24 = op2(e20, e23)) | (e24 = op2(e20, e22)) | (e24 = op2(e20, e21)) | (op2(e20, e20) = e24)) & ((e23 = op2(e26, e20)) | (e23 = op2(e25, e20)) | (e23 = op2(e24, e20)) | (e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e26)) | (e23 = op2(e20, e25)) | (e23 = op2(e20, e24)) | (e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e26, e20)) | (e22 = op2(e25, e20)) | (e22 = op2(e24, e20)) | (e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e26)) | (e22 = op2(e20, e25)) | (e22 = op2(e20, e24)) | (e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e26, e20)) | (e21 = op2(e25, e20)) | (e21 = op2(e24, e20)) | (e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e26)) | (e21 = op2(e20, e25)) | (e21 = op2(e20, e24)) | (e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e26, e20)) | (e20 = op2(e25, e20)) | (e20 = op2(e24, e20)) | (e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e26)) | (e20 = op2(e20, e25)) | (e20 = op2(e20, e24)) | (e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax5)).
fof(f4613, plain, (spl42_631 | spl42_624 | spl42_617 | spl42_610 | spl42_603 | spl42_596 | spl42_589), inference(avatar_split_clause, [], [f370, f4125, f4154, f4183, f4212, f4241, f4270, f4299])).
fof(f370, plain, ((e20 = op2(e21, e26)) | (e20 = op2(e21, e25)) | (e20 = op2(e21, e24)) | (e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f5])).
fof(f4604, plain, (spl42_677 | spl42_628 | spl42_579 | spl42_530 | spl42_481 | spl42_432 | spl42_383), inference(avatar_split_clause, [], [f379, f3271, f3474, f3677, f3880, f4083, f4286, f4489])).
fof(f379, plain, ((e24 = op2(e26, e21)) | (e24 = op2(e25, e21)) | (e24 = op2(e24, e21)) | (e24 = op2(e23, e21)) | (e24 = op2(e22, e21)) | (e24 = op2(e21, e21)) | (e24 = op2(e20, e21))), inference(cnf_transformation, [], [f5])).
fof(f4602, plain, (spl42_678 | spl42_629 | spl42_580 | spl42_531 | spl42_482 | spl42_433 | spl42_384), inference(avatar_split_clause, [], [f381, f3275, f3478, f3681, f3884, f4087, f4290, f4493])).
fof(f381, plain, ((e25 = op2(e26, e21)) | (e25 = op2(e25, e21)) | (e25 = op2(e24, e21)) | (e25 = op2(e23, e21)) | (e25 = op2(e22, e21)) | (e25 = op2(e21, e21)) | (e25 = op2(e20, e21))), inference(cnf_transformation, [], [f5])).
fof(f4601, plain, (spl42_637 | spl42_630 | spl42_623 | spl42_616 | spl42_609 | spl42_602 | spl42_595), inference(avatar_split_clause, [], [f382, f4149, f4178, f4207, f4236, f4265, f4294, f4323])).
fof(f382, plain, ((e26 = op2(e21, e26)) | (e26 = op2(e21, e25)) | (e26 = op2(e21, e24)) | (e26 = op2(e21, e23)) | (e26 = op2(e21, e22)) | (e26 = op2(e21, e21)) | (e26 = op2(e21, e20))), inference(cnf_transformation, [], [f5])).
fof(f4599, plain, (spl42_582 | spl42_575 | spl42_568 | spl42_561 | spl42_554 | spl42_547 | spl42_540), inference(avatar_split_clause, [], [f384, f3922, f3951, f3980, f4009, f4038, f4067, f4096])).
fof(f384, plain, ((e20 = op2(e22, e26)) | (e20 = op2(e22, e25)) | (e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f5])).
fof(f4586, plain, (spl42_672 | spl42_623 | spl42_574 | spl42_525 | spl42_476 | spl42_427 | spl42_378), inference(avatar_split_clause, [], [f397, f3250, f3453, f3656, f3859, f4062, f4265, f4468])).
fof(f397, plain, ((e26 = op2(e26, e22)) | (e26 = op2(e25, e22)) | (e26 = op2(e24, e22)) | (e26 = op2(e23, e22)) | (e26 = op2(e22, e22)) | (e26 = op2(e21, e22)) | (e26 = op2(e20, e22))), inference(cnf_transformation, [], [f5])).
fof(f4581, plain, (spl42_535 | spl42_528 | spl42_521 | spl42_514 | spl42_507 | spl42_500 | spl42_493), inference(avatar_split_clause, [], [f402, f3727, f3756, f3785, f3814, f3843, f3872, f3901])).
fof(f402, plain, ((e22 = op2(e23, e26)) | (e22 = op2(e23, e25)) | (e22 = op2(e23, e24)) | (e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))), inference(cnf_transformation, [], [f5])).
fof(f4580, plain, (spl42_661 | spl42_612 | spl42_563 | spl42_514 | spl42_465 | spl42_416 | spl42_367), inference(avatar_split_clause, [], [f403, f3205, f3408, f3611, f3814, f4017, f4220, f4423])).
fof(f403, plain, ((e22 = op2(e26, e23)) | (e22 = op2(e25, e23)) | (e22 = op2(e24, e23)) | (e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))), inference(cnf_transformation, [], [f5])).
fof(f4575, plain, (spl42_538 | spl42_531 | spl42_524 | spl42_517 | spl42_510 | spl42_503 | spl42_496), inference(avatar_split_clause, [], [f408, f3739, f3768, f3797, f3826, f3855, f3884, f3913])).
fof(f408, plain, ((e25 = op2(e23, e26)) | (e25 = op2(e23, e25)) | (e25 = op2(e23, e24)) | (e25 = op2(e23, e23)) | (e25 = op2(e23, e22)) | (e25 = op2(e23, e21)) | (e25 = op2(e23, e20))), inference(cnf_transformation, [], [f5])).
fof(f4574, plain, (spl42_664 | spl42_615 | spl42_566 | spl42_517 | spl42_468 | spl42_419 | spl42_370), inference(avatar_split_clause, [], [f409, f3217, f3420, f3623, f3826, f4029, f4232, f4435])).
fof(f409, plain, ((e25 = op2(e26, e23)) | (e25 = op2(e25, e23)) | (e25 = op2(e24, e23)) | (e25 = op2(e23, e23)) | (e25 = op2(e22, e23)) | (e25 = op2(e21, e23)) | (e25 = op2(e20, e23))), inference(cnf_transformation, [], [f5])).
fof(f4570, plain, (spl42_652 | spl42_603 | spl42_554 | spl42_505 | spl42_456 | spl42_407 | spl42_358), inference(avatar_split_clause, [], [f413, f3168, f3371, f3574, f3777, f3980, f4183, f4386])).
fof(f413, plain, ((e20 = op2(e26, e24)) | (e20 = op2(e25, e24)) | (e20 = op2(e24, e24)) | (e20 = op2(e23, e24)) | (e20 = op2(e22, e24)) | (e20 = op2(e21, e24)) | (e20 = op2(e20, e24))), inference(cnf_transformation, [], [f5])).
fof(f4560, plain, (spl42_657 | spl42_608 | spl42_559 | spl42_510 | spl42_461 | spl42_412 | spl42_363), inference(avatar_split_clause, [], [f423, f3188, f3391, f3594, f3797, f4000, f4203, f4406])).
fof(f423, plain, ((e25 = op2(e26, e24)) | (e25 = op2(e25, e24)) | (e25 = op2(e24, e24)) | (e25 = op2(e23, e24)) | (e25 = op2(e22, e24)) | (e25 = op2(e21, e24)) | (e25 = op2(e20, e24))), inference(cnf_transformation, [], [f5])).
fof(f4553, plain, (spl42_437 | spl42_430 | spl42_423 | spl42_416 | spl42_409 | spl42_402 | spl42_395), inference(avatar_split_clause, [], [f430, f3321, f3350, f3379, f3408, f3437, f3466, f3495])).
fof(f430, plain, ((e22 = op2(e25, e26)) | (e22 = op2(e25, e25)) | (e22 = op2(e25, e24)) | (e22 = op2(e25, e23)) | (e22 = op2(e25, e22)) | (e22 = op2(e25, e21)) | (e22 = op2(e25, e20))), inference(cnf_transformation, [], [f5])).
fof(f4549, plain, (spl42_439 | spl42_432 | spl42_425 | spl42_418 | spl42_411 | spl42_404 | spl42_397), inference(avatar_split_clause, [], [f434, f3329, f3358, f3387, f3416, f3445, f3474, f3503])).
fof(f434, plain, ((e24 = op2(e25, e26)) | (e24 = op2(e25, e25)) | (e24 = op2(e25, e24)) | (e24 = op2(e25, e23)) | (e24 = op2(e25, e22)) | (e24 = op2(e25, e21)) | (e24 = op2(e25, e20))), inference(cnf_transformation, [], [f5])).
fof(f4548, plain, (spl42_649 | spl42_600 | spl42_551 | spl42_502 | spl42_453 | spl42_404 | spl42_355), inference(avatar_split_clause, [], [f435, f3155, f3358, f3561, f3764, f3967, f4170, f4373])).
fof(f435, plain, ((e24 = op2(e26, e25)) | (e24 = op2(e25, e25)) | (e24 = op2(e24, e25)) | (e24 = op2(e23, e25)) | (e24 = op2(e22, e25)) | (e24 = op2(e21, e25)) | (e24 = op2(e20, e25))), inference(cnf_transformation, [], [f5])).
fof(f4547, plain, (spl42_440 | spl42_433 | spl42_426 | spl42_419 | spl42_412 | spl42_405 | spl42_398), inference(avatar_split_clause, [], [f436, f3333, f3362, f3391, f3420, f3449, f3478, f3507])).
fof(f436, plain, ((e25 = op2(e25, e26)) | (e25 = op2(e25, e25)) | (e25 = op2(e25, e24)) | (e25 = op2(e25, e23)) | (e25 = op2(e25, e22)) | (e25 = op2(e25, e21)) | (e25 = op2(e25, e20))), inference(cnf_transformation, [], [f5])).
fof(f4545, plain, (spl42_441 | spl42_434 | spl42_427 | spl42_420 | spl42_413 | spl42_406 | spl42_399), inference(avatar_split_clause, [], [f438, f3337, f3366, f3395, f3424, f3453, f3482, f3511])).
fof(f438, plain, ((e26 = op2(e25, e26)) | (e26 = op2(e25, e25)) | (e26 = op2(e25, e24)) | (e26 = op2(e25, e23)) | (e26 = op2(e25, e22)) | (e26 = op2(e25, e21)) | (e26 = op2(e25, e20))), inference(cnf_transformation, [], [f5])).
fof(f4544, plain, (spl42_651 | spl42_602 | spl42_553 | spl42_504 | spl42_455 | spl42_406 | spl42_357), inference(avatar_split_clause, [], [f439, f3163, f3366, f3569, f3772, f3975, f4178, f4381])).
fof(f439, plain, ((e26 = op2(e26, e25)) | (e26 = op2(e25, e25)) | (e26 = op2(e24, e25)) | (e26 = op2(e23, e25)) | (e26 = op2(e22, e25)) | (e26 = op2(e21, e25)) | (e26 = op2(e20, e25))), inference(cnf_transformation, [], [f5])).
fof(f4541, plain, (spl42_387 | spl42_380 | spl42_373 | spl42_366 | spl42_359 | spl42_352 | spl42_345), inference(avatar_split_clause, [], [f442, f3114, f3143, f3172, f3201, f3230, f3259, f3288])).
fof(f442, plain, ((e21 = op2(e26, e26)) | (e21 = op2(e26, e25)) | (e21 = op2(e26, e24)) | (e21 = op2(e26, e23)) | (e21 = op2(e26, e22)) | (e21 = op2(e26, e21)) | (e21 = op2(e26, e20))), inference(cnf_transformation, [], [f5])).
fof(f4540, plain, (spl42_639 | spl42_590 | spl42_541 | spl42_492 | spl42_443 | spl42_394 | spl42_345), inference(avatar_split_clause, [], [f443, f3114, f3317, f3520, f3723, f3926, f4129, f4332])).
fof(f443, plain, ((e21 = op2(e26, e26)) | (e21 = op2(e25, e26)) | (e21 = op2(e24, e26)) | (e21 = op2(e23, e26)) | (e21 = op2(e22, e26)) | (e21 = op2(e21, e26)) | (e21 = op2(e20, e26))), inference(cnf_transformation, [], [f5])).
fof(f4538, plain, (spl42_640 | spl42_591 | spl42_542 | spl42_493 | spl42_444 | spl42_395 | spl42_346), inference(avatar_split_clause, [], [f445, f3118, f3321, f3524, f3727, f3930, f4133, f4336])).
fof(f445, plain, ((e22 = op2(e26, e26)) | (e22 = op2(e25, e26)) | (e22 = op2(e24, e26)) | (e22 = op2(e23, e26)) | (e22 = op2(e22, e26)) | (e22 = op2(e21, e26)) | (e22 = op2(e20, e26))), inference(cnf_transformation, [], [f5])).
fof(f4534, plain, (spl42_642 | spl42_593 | spl42_544 | spl42_495 | spl42_446 | spl42_397 | spl42_348), inference(avatar_split_clause, [], [f449, f3126, f3329, f3532, f3735, f3938, f4141, f4344])).
fof(f449, plain, ((e24 = op2(e26, e26)) | (e24 = op2(e25, e26)) | (e24 = op2(e24, e26)) | (e24 = op2(e23, e26)) | (e24 = op2(e22, e26)) | (e24 = op2(e21, e26)) | (e24 = op2(e20, e26))), inference(cnf_transformation, [], [f5])).
fof(f4531, plain, (spl42_392 | spl42_385 | spl42_378 | spl42_371 | spl42_364 | spl42_357 | spl42_350), inference(avatar_split_clause, [], [f452, f3134, f3163, f3192, f3221, f3250, f3279, f3308])).
fof(f452, plain, ((e26 = op2(e26, e26)) | (e26 = op2(e26, e25)) | (e26 = op2(e26, e24)) | (e26 = op2(e26, e23)) | (e26 = op2(e26, e22)) | (e26 = op2(e26, e21)) | (e26 = op2(e26, e20))), inference(cnf_transformation, [], [f5])).
fof(f4500, plain, (spl42_673 | spl42_674 | spl42_675 | spl42_676 | spl42_677 | spl42_678 | spl42_679), inference(avatar_split_clause, [], [f308, f4497, f4493, f4489, f4485, f4481, f4477, f4473])).
fof(f308, plain, ((e26 = op2(e20, e21)) | (e25 = op2(e20, e21)) | (e24 = op2(e20, e21)) | (e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e26 = op2(e26, e26)) | (e25 = op2(e26, e26)) | (e24 = op2(e26, e26)) | (e23 = op2(e26, e26)) | (e22 = op2(e26, e26)) | (e21 = op2(e26, e26)) | (e20 = op2(e26, e26))) & ((e26 = op2(e26, e25)) | (e25 = op2(e26, e25)) | (e24 = op2(e26, e25)) | (e23 = op2(e26, e25)) | (e22 = op2(e26, e25)) | (e21 = op2(e26, e25)) | (e20 = op2(e26, e25))) & ((e26 = op2(e26, e24)) | (e25 = op2(e26, e24)) | (e24 = op2(e26, e24)) | (e23 = op2(e26, e24)) | (e22 = op2(e26, e24)) | (e21 = op2(e26, e24)) | (e20 = op2(e26, e24))) & ((e26 = op2(e26, e23)) | (e25 = op2(e26, e23)) | (e24 = op2(e26, e23)) | (e23 = op2(e26, e23)) | (e22 = op2(e26, e23)) | (e21 = op2(e26, e23)) | (e20 = op2(e26, e23))) & ((e26 = op2(e26, e22)) | (e25 = op2(e26, e22)) | (e24 = op2(e26, e22)) | (e23 = op2(e26, e22)) | (e22 = op2(e26, e22)) | (e21 = op2(e26, e22)) | (e20 = op2(e26, e22))) & ((e26 = op2(e26, e21)) | (e25 = op2(e26, e21)) | (e24 = op2(e26, e21)) | (e23 = op2(e26, e21)) | (e22 = op2(e26, e21)) | (e21 = op2(e26, e21)) | (e20 = op2(e26, e21))) & ((e26 = op2(e26, e20)) | (e25 = op2(e26, e20)) | (e24 = op2(e26, e20)) | (e23 = op2(e26, e20)) | (e22 = op2(e26, e20)) | (e21 = op2(e26, e20)) | (e20 = op2(e26, e20))) & ((e26 = op2(e25, e26)) | (e25 = op2(e25, e26)) | (e24 = op2(e25, e26)) | (e23 = op2(e25, e26)) | (e22 = op2(e25, e26)) | (e21 = op2(e25, e26)) | (e20 = op2(e25, e26))) & ((e26 = op2(e25, e25)) | (e25 = op2(e25, e25)) | (e24 = op2(e25, e25)) | (e23 = op2(e25, e25)) | (e22 = op2(e25, e25)) | (e21 = op2(e25, e25)) | (e20 = op2(e25, e25))) & ((e26 = op2(e25, e24)) | (e25 = op2(e25, e24)) | (e24 = op2(e25, e24)) | (e23 = op2(e25, e24)) | (e22 = op2(e25, e24)) | (e21 = op2(e25, e24)) | (e20 = op2(e25, e24))) & ((e26 = op2(e25, e23)) | (e25 = op2(e25, e23)) | (e24 = op2(e25, e23)) | (e23 = op2(e25, e23)) | (e22 = op2(e25, e23)) | (e21 = op2(e25, e23)) | (e20 = op2(e25, e23))) & ((e26 = op2(e25, e22)) | (e25 = op2(e25, e22)) | (e24 = op2(e25, e22)) | (e23 = op2(e25, e22)) | (e22 = op2(e25, e22)) | (e21 = op2(e25, e22)) | (e20 = op2(e25, e22))) & ((e26 = op2(e25, e21)) | (e25 = op2(e25, e21)) | (e24 = op2(e25, e21)) | (e23 = op2(e25, e21)) | (e22 = op2(e25, e21)) | (e21 = op2(e25, e21)) | (e20 = op2(e25, e21))) & ((e26 = op2(e25, e20)) | (e25 = op2(e25, e20)) | (e24 = op2(e25, e20)) | (e23 = op2(e25, e20)) | (e22 = op2(e25, e20)) | (e21 = op2(e25, e20)) | (e20 = op2(e25, e20))) & ((e26 = op2(e24, e26)) | (e25 = op2(e24, e26)) | (e24 = op2(e24, e26)) | (e23 = op2(e24, e26)) | (e22 = op2(e24, e26)) | (e21 = op2(e24, e26)) | (e20 = op2(e24, e26))) & ((e26 = op2(e24, e25)) | (e25 = op2(e24, e25)) | (e24 = op2(e24, e25)) | (e23 = op2(e24, e25)) | (e22 = op2(e24, e25)) | (e21 = op2(e24, e25)) | (e20 = op2(e24, e25))) & ((e26 = op2(e24, e24)) | (e25 = op2(e24, e24)) | (e24 = op2(e24, e24)) | (e23 = op2(e24, e24)) | (e22 = op2(e24, e24)) | (e21 = op2(e24, e24)) | (e20 = op2(e24, e24))) & ((e26 = op2(e24, e23)) | (e25 = op2(e24, e23)) | (e24 = op2(e24, e23)) | (e23 = op2(e24, e23)) | (e22 = op2(e24, e23)) | (e21 = op2(e24, e23)) | (e20 = op2(e24, e23))) & ((e26 = op2(e24, e22)) | (e25 = op2(e24, e22)) | (e24 = op2(e24, e22)) | (e23 = op2(e24, e22)) | (e22 = op2(e24, e22)) | (e21 = op2(e24, e22)) | (e20 = op2(e24, e22))) & ((e26 = op2(e24, e21)) | (e25 = op2(e24, e21)) | (e24 = op2(e24, e21)) | (e23 = op2(e24, e21)) | (e22 = op2(e24, e21)) | (e21 = op2(e24, e21)) | (e20 = op2(e24, e21))) & ((e26 = op2(e24, e20)) | (e25 = op2(e24, e20)) | (e24 = op2(e24, e20)) | (e23 = op2(e24, e20)) | (e22 = op2(e24, e20)) | (e21 = op2(e24, e20)) | (e20 = op2(e24, e20))) & ((e26 = op2(e23, e26)) | (e25 = op2(e23, e26)) | (e24 = op2(e23, e26)) | (e23 = op2(e23, e26)) | (e22 = op2(e23, e26)) | (e21 = op2(e23, e26)) | (e20 = op2(e23, e26))) & ((e26 = op2(e23, e25)) | (e25 = op2(e23, e25)) | (e24 = op2(e23, e25)) | (e23 = op2(e23, e25)) | (e22 = op2(e23, e25)) | (e21 = op2(e23, e25)) | (e20 = op2(e23, e25))) & ((e26 = op2(e23, e24)) | (e25 = op2(e23, e24)) | (e24 = op2(e23, e24)) | (e23 = op2(e23, e24)) | (e22 = op2(e23, e24)) | (e21 = op2(e23, e24)) | (e20 = op2(e23, e24))) & ((e26 = op2(e23, e23)) | (e25 = op2(e23, e23)) | (e24 = op2(e23, e23)) | (e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e26 = op2(e23, e22)) | (e25 = op2(e23, e22)) | (e24 = op2(e23, e22)) | (e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e26 = op2(e23, e21)) | (e25 = op2(e23, e21)) | (e24 = op2(e23, e21)) | (e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e26 = op2(e23, e20)) | (e25 = op2(e23, e20)) | (e24 = op2(e23, e20)) | (e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e26 = op2(e22, e26)) | (e25 = op2(e22, e26)) | (e24 = op2(e22, e26)) | (e23 = op2(e22, e26)) | (e22 = op2(e22, e26)) | (e21 = op2(e22, e26)) | (e20 = op2(e22, e26))) & ((e26 = op2(e22, e25)) | (e25 = op2(e22, e25)) | (e24 = op2(e22, e25)) | (e23 = op2(e22, e25)) | (e22 = op2(e22, e25)) | (e21 = op2(e22, e25)) | (e20 = op2(e22, e25))) & ((e26 = op2(e22, e24)) | (e25 = op2(e22, e24)) | (e24 = op2(e22, e24)) | (e23 = op2(e22, e24)) | (e22 = op2(e22, e24)) | (e21 = op2(e22, e24)) | (e20 = op2(e22, e24))) & ((e26 = op2(e22, e23)) | (e25 = op2(e22, e23)) | (e24 = op2(e22, e23)) | (e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e26 = op2(e22, e22)) | (e25 = op2(e22, e22)) | (e24 = op2(e22, e22)) | (e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e26 = op2(e22, e21)) | (e25 = op2(e22, e21)) | (e24 = op2(e22, e21)) | (e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e26 = op2(e22, e20)) | (e25 = op2(e22, e20)) | (e24 = op2(e22, e20)) | (e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e26 = op2(e21, e26)) | (e25 = op2(e21, e26)) | (e24 = op2(e21, e26)) | (e23 = op2(e21, e26)) | (e22 = op2(e21, e26)) | (e21 = op2(e21, e26)) | (e20 = op2(e21, e26))) & ((e26 = op2(e21, e25)) | (e25 = op2(e21, e25)) | (e24 = op2(e21, e25)) | (e23 = op2(e21, e25)) | (e22 = op2(e21, e25)) | (e21 = op2(e21, e25)) | (e20 = op2(e21, e25))) & ((e26 = op2(e21, e24)) | (e25 = op2(e21, e24)) | (e24 = op2(e21, e24)) | (e23 = op2(e21, e24)) | (e22 = op2(e21, e24)) | (e21 = op2(e21, e24)) | (e20 = op2(e21, e24))) & ((e26 = op2(e21, e23)) | (e25 = op2(e21, e23)) | (e24 = op2(e21, e23)) | (e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e26 = op2(e21, e22)) | (e25 = op2(e21, e22)) | (e24 = op2(e21, e22)) | (e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e26 = op2(e21, e21)) | (e25 = op2(e21, e21)) | (e24 = op2(e21, e21)) | (e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e26 = op2(e21, e20)) | (e25 = op2(e21, e20)) | (e24 = op2(e21, e20)) | (e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e26 = op2(e20, e26)) | (e25 = op2(e20, e26)) | (e24 = op2(e20, e26)) | (e23 = op2(e20, e26)) | (e22 = op2(e20, e26)) | (e21 = op2(e20, e26)) | (e20 = op2(e20, e26))) & ((e26 = op2(e20, e25)) | (e25 = op2(e20, e25)) | (e24 = op2(e20, e25)) | (e23 = op2(e20, e25)) | (e22 = op2(e20, e25)) | (e21 = op2(e20, e25)) | (e20 = op2(e20, e25))) & ((e26 = op2(e20, e24)) | (e25 = op2(e20, e24)) | (e24 = op2(e20, e24)) | (e23 = op2(e20, e24)) | (e22 = op2(e20, e24)) | (e21 = op2(e20, e24)) | (e20 = op2(e20, e24))) & ((e26 = op2(e20, e23)) | (e25 = op2(e20, e23)) | (e24 = op2(e20, e23)) | (e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e26 = op2(e20, e22)) | (e25 = op2(e20, e22)) | (e24 = op2(e20, e22)) | (e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e26 = op2(e20, e21)) | (e25 = op2(e20, e21)) | (e24 = op2(e20, e21)) | (e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e26) | (op2(e20, e20) = e25) | (op2(e20, e20) = e24) | (op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax4)).
fof(f4471, plain, (spl42_666 | spl42_667 | spl42_668 | spl42_669 | spl42_670 | spl42_671 | spl42_672), inference(avatar_split_clause, [], [f309, f4468, f4464, f4460, f4456, f4452, f4448, f4444])).
fof(f309, plain, ((e26 = op2(e20, e22)) | (e25 = op2(e20, e22)) | (e24 = op2(e20, e22)) | (e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f4413, plain, (spl42_652 | spl42_653 | spl42_654 | spl42_655 | spl42_656 | spl42_657 | spl42_658), inference(avatar_split_clause, [], [f311, f4410, f4406, f4402, f4398, f4394, f4390, f4386])).
fof(f311, plain, ((e26 = op2(e20, e24)) | (e25 = op2(e20, e24)) | (e24 = op2(e20, e24)) | (e23 = op2(e20, e24)) | (e22 = op2(e20, e24)) | (e21 = op2(e20, e24)) | (e20 = op2(e20, e24))), inference(cnf_transformation, [], [f4])).
fof(f4268, plain, (spl42_617 | spl42_618 | spl42_619 | spl42_620 | spl42_621 | spl42_622 | spl42_623), inference(avatar_split_clause, [], [f316, f4265, f4261, f4257, f4253, f4249, f4245, f4241])).
fof(f316, plain, ((e26 = op2(e21, e22)) | (e25 = op2(e21, e22)) | (e24 = op2(e21, e22)) | (e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))), inference(cnf_transformation, [], [f4])).
fof(f4239, plain, (spl42_610 | spl42_611 | spl42_612 | spl42_613 | spl42_614 | spl42_615 | spl42_616), inference(avatar_split_clause, [], [f317, f4236, f4232, f4228, f4224, f4220, f4216, f4212])).
fof(f317, plain, ((e26 = op2(e21, e23)) | (e25 = op2(e21, e23)) | (e24 = op2(e21, e23)) | (e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))), inference(cnf_transformation, [], [f4])).
fof(f4210, plain, (spl42_603 | spl42_604 | spl42_605 | spl42_606 | spl42_607 | spl42_608 | spl42_609), inference(avatar_split_clause, [], [f318, f4207, f4203, f4199, f4195, f4191, f4187, f4183])).
fof(f318, plain, ((e26 = op2(e21, e24)) | (e25 = op2(e21, e24)) | (e24 = op2(e21, e24)) | (e23 = op2(e21, e24)) | (e22 = op2(e21, e24)) | (e21 = op2(e21, e24)) | (e20 = op2(e21, e24))), inference(cnf_transformation, [], [f4])).
fof(f4152, plain, (spl42_589 | spl42_590 | spl42_591 | spl42_592 | spl42_593 | spl42_594 | spl42_595), inference(avatar_split_clause, [], [f320, f4149, f4145, f4141, f4137, f4133, f4129, f4125])).
fof(f320, plain, ((e26 = op2(e21, e26)) | (e25 = op2(e21, e26)) | (e24 = op2(e21, e26)) | (e23 = op2(e21, e26)) | (e22 = op2(e21, e26)) | (e21 = op2(e21, e26)) | (e20 = op2(e21, e26))), inference(cnf_transformation, [], [f4])).
fof(f4123, plain, (spl42_582 | spl42_583 | spl42_584 | spl42_585 | spl42_586 | spl42_587 | spl42_588), inference(avatar_split_clause, [], [f321, f4120, f4116, f4112, f4108, f4104, f4100, f4096])).
fof(f321, plain, ((e26 = op2(e22, e20)) | (e25 = op2(e22, e20)) | (e24 = op2(e22, e20)) | (e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f3949, plain, (spl42_540 | spl42_541 | spl42_542 | spl42_543 | spl42_544 | spl42_545 | spl42_546), inference(avatar_split_clause, [], [f327, f3946, f3942, f3938, f3934, f3930, f3926, f3922])).
fof(f327, plain, ((e26 = op2(e22, e26)) | (e25 = op2(e22, e26)) | (e24 = op2(e22, e26)) | (e23 = op2(e22, e26)) | (e22 = op2(e22, e26)) | (e21 = op2(e22, e26)) | (e20 = op2(e22, e26))), inference(cnf_transformation, [], [f4])).
fof(f3891, plain, (spl42_526 | spl42_527 | spl42_528 | spl42_529 | spl42_530 | spl42_531 | spl42_532), inference(avatar_split_clause, [], [f329, f3888, f3884, f3880, f3876, f3872, f3868, f3864])).
fof(f329, plain, ((e26 = op2(e23, e21)) | (e25 = op2(e23, e21)) | (e24 = op2(e23, e21)) | (e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))), inference(cnf_transformation, [], [f4])).
fof(f3833, plain, (spl42_512 | spl42_513 | spl42_514 | spl42_515 | spl42_516 | spl42_517 | spl42_518), inference(avatar_split_clause, [], [f331, f3830, f3826, f3822, f3818, f3814, f3810, f3806])).
fof(f331, plain, ((e26 = op2(e23, e23)) | (e25 = op2(e23, e23)) | (e24 = op2(e23, e23)) | (e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))), inference(cnf_transformation, [], [f4])).
fof(f3775, plain, (spl42_498 | spl42_499 | spl42_500 | spl42_501 | spl42_502 | spl42_503 | spl42_504), inference(avatar_split_clause, [], [f333, f3772, f3768, f3764, f3760, f3756, f3752, f3748])).
fof(f333, plain, ((e26 = op2(e23, e25)) | (e25 = op2(e23, e25)) | (e24 = op2(e23, e25)) | (e23 = op2(e23, e25)) | (e22 = op2(e23, e25)) | (e21 = op2(e23, e25)) | (e20 = op2(e23, e25))), inference(cnf_transformation, [], [f4])).
fof(f3746, plain, (spl42_491 | spl42_492 | spl42_493 | spl42_494 | spl42_495 | spl42_496 | spl42_497), inference(avatar_split_clause, [], [f334, f3743, f3739, f3735, f3731, f3727, f3723, f3719])).
fof(f334, plain, ((e26 = op2(e23, e26)) | (e25 = op2(e23, e26)) | (e24 = op2(e23, e26)) | (e23 = op2(e23, e26)) | (e22 = op2(e23, e26)) | (e21 = op2(e23, e26)) | (e20 = op2(e23, e26))), inference(cnf_transformation, [], [f4])).
fof(f3543, plain, (spl42_442 | spl42_443 | spl42_444 | spl42_445 | spl42_446 | spl42_447 | spl42_448), inference(avatar_split_clause, [], [f341, f3540, f3536, f3532, f3528, f3524, f3520, f3516])).
fof(f341, plain, ((e26 = op2(e24, e26)) | (e25 = op2(e24, e26)) | (e24 = op2(e24, e26)) | (e23 = op2(e24, e26)) | (e22 = op2(e24, e26)) | (e21 = op2(e24, e26)) | (e20 = op2(e24, e26))), inference(cnf_transformation, [], [f4])).
fof(f3514, plain, (spl42_435 | spl42_436 | spl42_437 | spl42_438 | spl42_439 | spl42_440 | spl42_441), inference(avatar_split_clause, [], [f342, f3511, f3507, f3503, f3499, f3495, f3491, f3487])).
fof(f342, plain, ((e26 = op2(e25, e20)) | (e25 = op2(e25, e20)) | (e24 = op2(e25, e20)) | (e23 = op2(e25, e20)) | (e22 = op2(e25, e20)) | (e21 = op2(e25, e20)) | (e20 = op2(e25, e20))), inference(cnf_transformation, [], [f4])).
fof(f3427, plain, (spl42_414 | spl42_415 | spl42_416 | spl42_417 | spl42_418 | spl42_419 | spl42_420), inference(avatar_split_clause, [], [f345, f3424, f3420, f3416, f3412, f3408, f3404, f3400])).
fof(f345, plain, ((e26 = op2(e25, e23)) | (e25 = op2(e25, e23)) | (e24 = op2(e25, e23)) | (e23 = op2(e25, e23)) | (e22 = op2(e25, e23)) | (e21 = op2(e25, e23)) | (e20 = op2(e25, e23))), inference(cnf_transformation, [], [f4])).
fof(f3340, plain, (spl42_393 | spl42_394 | spl42_395 | spl42_396 | spl42_397 | spl42_398 | spl42_399), inference(avatar_split_clause, [], [f348, f3337, f3333, f3329, f3325, f3321, f3317, f3313])).
fof(f348, plain, ((e26 = op2(e25, e26)) | (e25 = op2(e25, e26)) | (e24 = op2(e25, e26)) | (e23 = op2(e25, e26)) | (e22 = op2(e25, e26)) | (e21 = op2(e25, e26)) | (e20 = op2(e25, e26))), inference(cnf_transformation, [], [f4])).
fof(f3311, plain, (spl42_386 | spl42_387 | spl42_388 | spl42_389 | spl42_390 | spl42_391 | spl42_392), inference(avatar_split_clause, [], [f349, f3308, f3304, f3300, f3296, f3292, f3288, f3284])).
fof(f349, plain, ((e26 = op2(e26, e20)) | (e25 = op2(e26, e20)) | (e24 = op2(e26, e20)) | (e23 = op2(e26, e20)) | (e22 = op2(e26, e20)) | (e21 = op2(e26, e20)) | (e20 = op2(e26, e20))), inference(cnf_transformation, [], [f4])).
fof(f3253, plain, (spl42_372 | spl42_373 | spl42_374 | spl42_375 | spl42_376 | spl42_377 | spl42_378), inference(avatar_split_clause, [], [f351, f3250, f3246, f3242, f3238, f3234, f3230, f3226])).
fof(f351, plain, ((e26 = op2(e26, e22)) | (e25 = op2(e26, e22)) | (e24 = op2(e26, e22)) | (e23 = op2(e26, e22)) | (e22 = op2(e26, e22)) | (e21 = op2(e26, e22)) | (e20 = op2(e26, e22))), inference(cnf_transformation, [], [f4])).
fof(f3195, plain, (spl42_358 | spl42_359 | spl42_360 | spl42_361 | spl42_362 | spl42_363 | spl42_364), inference(avatar_split_clause, [], [f353, f3192, f3188, f3184, f3180, f3176, f3172, f3168])).
fof(f353, plain, ((e26 = op2(e26, e24)) | (e25 = op2(e26, e24)) | (e24 = op2(e26, e24)) | (e23 = op2(e26, e24)) | (e22 = op2(e26, e24)) | (e21 = op2(e26, e24)) | (e20 = op2(e26, e24))), inference(cnf_transformation, [], [f4])).
fof(f3101, plain, (spl42_340 | spl42_291 | spl42_242 | spl42_193 | spl42_144 | spl42_95 | spl42_46), inference(avatar_split_clause, [], [f167, f1777, f1980, f2183, f2386, f2589, f2792, f2995])).
fof(f167, plain, ((e13 = op1(e16, e10)) | (e13 = op1(e15, e10)) | (e13 = op1(e14, e10)) | (e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e16 = op1(e16, e16)) | (e16 = op1(e15, e16)) | (e16 = op1(e14, e16)) | (e16 = op1(e13, e16)) | (e16 = op1(e12, e16)) | (e16 = op1(e11, e16)) | (e16 = op1(e10, e16))) & ((e16 = op1(e16, e16)) | (e16 = op1(e16, e15)) | (e16 = op1(e16, e14)) | (e16 = op1(e16, e13)) | (e16 = op1(e16, e12)) | (e16 = op1(e16, e11)) | (e16 = op1(e16, e10))) & ((e15 = op1(e16, e16)) | (e15 = op1(e15, e16)) | (e15 = op1(e14, e16)) | (e15 = op1(e13, e16)) | (e15 = op1(e12, e16)) | (e15 = op1(e11, e16)) | (e15 = op1(e10, e16))) & ((e15 = op1(e16, e16)) | (e15 = op1(e16, e15)) | (e15 = op1(e16, e14)) | (e15 = op1(e16, e13)) | (e15 = op1(e16, e12)) | (e15 = op1(e16, e11)) | (e15 = op1(e16, e10))) & ((e14 = op1(e16, e16)) | (e14 = op1(e15, e16)) | (e14 = op1(e14, e16)) | (e14 = op1(e13, e16)) | (e14 = op1(e12, e16)) | (e14 = op1(e11, e16)) | (e14 = op1(e10, e16))) & ((e14 = op1(e16, e16)) | (e14 = op1(e16, e15)) | (e14 = op1(e16, e14)) | (e14 = op1(e16, e13)) | (e14 = op1(e16, e12)) | (e14 = op1(e16, e11)) | (e14 = op1(e16, e10))) & ((e13 = op1(e16, e16)) | (e13 = op1(e15, e16)) | (e13 = op1(e14, e16)) | (e13 = op1(e13, e16)) | (e13 = op1(e12, e16)) | (e13 = op1(e11, e16)) | (e13 = op1(e10, e16))) & ((e13 = op1(e16, e16)) | (e13 = op1(e16, e15)) | (e13 = op1(e16, e14)) | (e13 = op1(e16, e13)) | (e13 = op1(e16, e12)) | (e13 = op1(e16, e11)) | (e13 = op1(e16, e10))) & ((e12 = op1(e16, e16)) | (e12 = op1(e15, e16)) | (e12 = op1(e14, e16)) | (e12 = op1(e13, e16)) | (e12 = op1(e12, e16)) | (e12 = op1(e11, e16)) | (e12 = op1(e10, e16))) & ((e12 = op1(e16, e16)) | (e12 = op1(e16, e15)) | (e12 = op1(e16, e14)) | (e12 = op1(e16, e13)) | (e12 = op1(e16, e12)) | (e12 = op1(e16, e11)) | (e12 = op1(e16, e10))) & ((e11 = op1(e16, e16)) | (e11 = op1(e15, e16)) | (e11 = op1(e14, e16)) | (e11 = op1(e13, e16)) | (e11 = op1(e12, e16)) | (e11 = op1(e11, e16)) | (e11 = op1(e10, e16))) & ((e11 = op1(e16, e16)) | (e11 = op1(e16, e15)) | (e11 = op1(e16, e14)) | (e11 = op1(e16, e13)) | (e11 = op1(e16, e12)) | (e11 = op1(e16, e11)) | (e11 = op1(e16, e10))) & ((e10 = op1(e16, e16)) | (e10 = op1(e15, e16)) | (e10 = op1(e14, e16)) | (e10 = op1(e13, e16)) | (e10 = op1(e12, e16)) | (e10 = op1(e11, e16)) | (e10 = op1(e10, e16))) & ((e10 = op1(e16, e16)) | (e10 = op1(e16, e15)) | (e10 = op1(e16, e14)) | (e10 = op1(e16, e13)) | (e10 = op1(e16, e12)) | (e10 = op1(e16, e11)) | (e10 = op1(e16, e10))) & ((e16 = op1(e16, e15)) | (e16 = op1(e15, e15)) | (e16 = op1(e14, e15)) | (e16 = op1(e13, e15)) | (e16 = op1(e12, e15)) | (e16 = op1(e11, e15)) | (e16 = op1(e10, e15))) & ((e16 = op1(e15, e16)) | (e16 = op1(e15, e15)) | (e16 = op1(e15, e14)) | (e16 = op1(e15, e13)) | (e16 = op1(e15, e12)) | (e16 = op1(e15, e11)) | (e16 = op1(e15, e10))) & ((e15 = op1(e16, e15)) | (e15 = op1(e15, e15)) | (e15 = op1(e14, e15)) | (e15 = op1(e13, e15)) | (e15 = op1(e12, e15)) | (e15 = op1(e11, e15)) | (e15 = op1(e10, e15))) & ((e15 = op1(e15, e16)) | (e15 = op1(e15, e15)) | (e15 = op1(e15, e14)) | (e15 = op1(e15, e13)) | (e15 = op1(e15, e12)) | (e15 = op1(e15, e11)) | (e15 = op1(e15, e10))) & ((e14 = op1(e16, e15)) | (e14 = op1(e15, e15)) | (e14 = op1(e14, e15)) | (e14 = op1(e13, e15)) | (e14 = op1(e12, e15)) | (e14 = op1(e11, e15)) | (e14 = op1(e10, e15))) & ((e14 = op1(e15, e16)) | (e14 = op1(e15, e15)) | (e14 = op1(e15, e14)) | (e14 = op1(e15, e13)) | (e14 = op1(e15, e12)) | (e14 = op1(e15, e11)) | (e14 = op1(e15, e10))) & ((e13 = op1(e16, e15)) | (e13 = op1(e15, e15)) | (e13 = op1(e14, e15)) | (e13 = op1(e13, e15)) | (e13 = op1(e12, e15)) | (e13 = op1(e11, e15)) | (e13 = op1(e10, e15))) & ((e13 = op1(e15, e16)) | (e13 = op1(e15, e15)) | (e13 = op1(e15, e14)) | (e13 = op1(e15, e13)) | (e13 = op1(e15, e12)) | (e13 = op1(e15, e11)) | (e13 = op1(e15, e10))) & ((e12 = op1(e16, e15)) | (e12 = op1(e15, e15)) | (e12 = op1(e14, e15)) | (e12 = op1(e13, e15)) | (e12 = op1(e12, e15)) | (e12 = op1(e11, e15)) | (e12 = op1(e10, e15))) & ((e12 = op1(e15, e16)) | (e12 = op1(e15, e15)) | (e12 = op1(e15, e14)) | (e12 = op1(e15, e13)) | (e12 = op1(e15, e12)) | (e12 = op1(e15, e11)) | (e12 = op1(e15, e10))) & ((e11 = op1(e16, e15)) | (e11 = op1(e15, e15)) | (e11 = op1(e14, e15)) | (e11 = op1(e13, e15)) | (e11 = op1(e12, e15)) | (e11 = op1(e11, e15)) | (e11 = op1(e10, e15))) & ((e11 = op1(e15, e16)) | (e11 = op1(e15, e15)) | (e11 = op1(e15, e14)) | (e11 = op1(e15, e13)) | (e11 = op1(e15, e12)) | (e11 = op1(e15, e11)) | (e11 = op1(e15, e10))) & ((e10 = op1(e16, e15)) | (e10 = op1(e15, e15)) | (e10 = op1(e14, e15)) | (e10 = op1(e13, e15)) | (e10 = op1(e12, e15)) | (e10 = op1(e11, e15)) | (e10 = op1(e10, e15))) & ((e10 = op1(e15, e16)) | (e10 = op1(e15, e15)) | (e10 = op1(e15, e14)) | (e10 = op1(e15, e13)) | (e10 = op1(e15, e12)) | (e10 = op1(e15, e11)) | (e10 = op1(e15, e10))) & ((e16 = op1(e16, e14)) | (e16 = op1(e15, e14)) | (e16 = op1(e14, e14)) | (e16 = op1(e13, e14)) | (e16 = op1(e12, e14)) | (e16 = op1(e11, e14)) | (e16 = op1(e10, e14))) & ((e16 = op1(e14, e16)) | (e16 = op1(e14, e15)) | (e16 = op1(e14, e14)) | (e16 = op1(e14, e13)) | (e16 = op1(e14, e12)) | (e16 = op1(e14, e11)) | (e16 = op1(e14, e10))) & ((e15 = op1(e16, e14)) | (e15 = op1(e15, e14)) | (e15 = op1(e14, e14)) | (e15 = op1(e13, e14)) | (e15 = op1(e12, e14)) | (e15 = op1(e11, e14)) | (e15 = op1(e10, e14))) & ((e15 = op1(e14, e16)) | (e15 = op1(e14, e15)) | (e15 = op1(e14, e14)) | (e15 = op1(e14, e13)) | (e15 = op1(e14, e12)) | (e15 = op1(e14, e11)) | (e15 = op1(e14, e10))) & ((e14 = op1(e16, e14)) | (e14 = op1(e15, e14)) | (e14 = op1(e14, e14)) | (e14 = op1(e13, e14)) | (e14 = op1(e12, e14)) | (e14 = op1(e11, e14)) | (e14 = op1(e10, e14))) & ((e14 = op1(e14, e16)) | (e14 = op1(e14, e15)) | (e14 = op1(e14, e14)) | (e14 = op1(e14, e13)) | (e14 = op1(e14, e12)) | (e14 = op1(e14, e11)) | (e14 = op1(e14, e10))) & ((e13 = op1(e16, e14)) | (e13 = op1(e15, e14)) | (e13 = op1(e14, e14)) | (e13 = op1(e13, e14)) | (e13 = op1(e12, e14)) | (e13 = op1(e11, e14)) | (e13 = op1(e10, e14))) & ((e13 = op1(e14, e16)) | (e13 = op1(e14, e15)) | (e13 = op1(e14, e14)) | (e13 = op1(e14, e13)) | (e13 = op1(e14, e12)) | (e13 = op1(e14, e11)) | (e13 = op1(e14, e10))) & ((e12 = op1(e16, e14)) | (e12 = op1(e15, e14)) | (e12 = op1(e14, e14)) | (e12 = op1(e13, e14)) | (e12 = op1(e12, e14)) | (e12 = op1(e11, e14)) | (e12 = op1(e10, e14))) & ((e12 = op1(e14, e16)) | (e12 = op1(e14, e15)) | (e12 = op1(e14, e14)) | (e12 = op1(e14, e13)) | (e12 = op1(e14, e12)) | (e12 = op1(e14, e11)) | (e12 = op1(e14, e10))) & ((e11 = op1(e16, e14)) | (e11 = op1(e15, e14)) | (e11 = op1(e14, e14)) | (e11 = op1(e13, e14)) | (e11 = op1(e12, e14)) | (e11 = op1(e11, e14)) | (e11 = op1(e10, e14))) & ((e11 = op1(e14, e16)) | (e11 = op1(e14, e15)) | (e11 = op1(e14, e14)) | (e11 = op1(e14, e13)) | (e11 = op1(e14, e12)) | (e11 = op1(e14, e11)) | (e11 = op1(e14, e10))) & ((e10 = op1(e16, e14)) | (e10 = op1(e15, e14)) | (e10 = op1(e14, e14)) | (e10 = op1(e13, e14)) | (e10 = op1(e12, e14)) | (e10 = op1(e11, e14)) | (e10 = op1(e10, e14))) & ((e10 = op1(e14, e16)) | (e10 = op1(e14, e15)) | (e10 = op1(e14, e14)) | (e10 = op1(e14, e13)) | (e10 = op1(e14, e12)) | (e10 = op1(e14, e11)) | (e10 = op1(e14, e10))) & ((e16 = op1(e16, e13)) | (e16 = op1(e15, e13)) | (e16 = op1(e14, e13)) | (e16 = op1(e13, e13)) | (e16 = op1(e12, e13)) | (e16 = op1(e11, e13)) | (e16 = op1(e10, e13))) & ((e16 = op1(e13, e16)) | (e16 = op1(e13, e15)) | (e16 = op1(e13, e14)) | (e16 = op1(e13, e13)) | (e16 = op1(e13, e12)) | (e16 = op1(e13, e11)) | (e16 = op1(e13, e10))) & ((e15 = op1(e16, e13)) | (e15 = op1(e15, e13)) | (e15 = op1(e14, e13)) | (e15 = op1(e13, e13)) | (e15 = op1(e12, e13)) | (e15 = op1(e11, e13)) | (e15 = op1(e10, e13))) & ((e15 = op1(e13, e16)) | (e15 = op1(e13, e15)) | (e15 = op1(e13, e14)) | (e15 = op1(e13, e13)) | (e15 = op1(e13, e12)) | (e15 = op1(e13, e11)) | (e15 = op1(e13, e10))) & ((e14 = op1(e16, e13)) | (e14 = op1(e15, e13)) | (e14 = op1(e14, e13)) | (e14 = op1(e13, e13)) | (e14 = op1(e12, e13)) | (e14 = op1(e11, e13)) | (e14 = op1(e10, e13))) & ((e14 = op1(e13, e16)) | (e14 = op1(e13, e15)) | (e14 = op1(e13, e14)) | (e14 = op1(e13, e13)) | (e14 = op1(e13, e12)) | (e14 = op1(e13, e11)) | (e14 = op1(e13, e10))) & ((e13 = op1(e16, e13)) | (e13 = op1(e15, e13)) | (e13 = op1(e14, e13)) | (e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e16)) | (e13 = op1(e13, e15)) | (e13 = op1(e13, e14)) | (e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e16, e13)) | (e12 = op1(e15, e13)) | (e12 = op1(e14, e13)) | (e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e16)) | (e12 = op1(e13, e15)) | (e12 = op1(e13, e14)) | (e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e16, e13)) | (e11 = op1(e15, e13)) | (e11 = op1(e14, e13)) | (e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e16)) | (e11 = op1(e13, e15)) | (e11 = op1(e13, e14)) | (e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e16, e13)) | (e10 = op1(e15, e13)) | (e10 = op1(e14, e13)) | (e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e16)) | (e10 = op1(e13, e15)) | (e10 = op1(e13, e14)) | (e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e16 = op1(e16, e12)) | (e16 = op1(e15, e12)) | (e16 = op1(e14, e12)) | (e16 = op1(e13, e12)) | (e16 = op1(e12, e12)) | (e16 = op1(e11, e12)) | (e16 = op1(e10, e12))) & ((e16 = op1(e12, e16)) | (e16 = op1(e12, e15)) | (e16 = op1(e12, e14)) | (e16 = op1(e12, e13)) | (e16 = op1(e12, e12)) | (e16 = op1(e12, e11)) | (e16 = op1(e12, e10))) & ((e15 = op1(e16, e12)) | (e15 = op1(e15, e12)) | (e15 = op1(e14, e12)) | (e15 = op1(e13, e12)) | (e15 = op1(e12, e12)) | (e15 = op1(e11, e12)) | (e15 = op1(e10, e12))) & ((e15 = op1(e12, e16)) | (e15 = op1(e12, e15)) | (e15 = op1(e12, e14)) | (e15 = op1(e12, e13)) | (e15 = op1(e12, e12)) | (e15 = op1(e12, e11)) | (e15 = op1(e12, e10))) & ((e14 = op1(e16, e12)) | (e14 = op1(e15, e12)) | (e14 = op1(e14, e12)) | (e14 = op1(e13, e12)) | (e14 = op1(e12, e12)) | (e14 = op1(e11, e12)) | (e14 = op1(e10, e12))) & ((e14 = op1(e12, e16)) | (e14 = op1(e12, e15)) | (e14 = op1(e12, e14)) | (e14 = op1(e12, e13)) | (e14 = op1(e12, e12)) | (e14 = op1(e12, e11)) | (e14 = op1(e12, e10))) & ((e13 = op1(e16, e12)) | (e13 = op1(e15, e12)) | (e13 = op1(e14, e12)) | (e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e16)) | (e13 = op1(e12, e15)) | (e13 = op1(e12, e14)) | (e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e16, e12)) | (e12 = op1(e15, e12)) | (e12 = op1(e14, e12)) | (e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e16)) | (e12 = op1(e12, e15)) | (e12 = op1(e12, e14)) | (e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e16, e12)) | (e11 = op1(e15, e12)) | (e11 = op1(e14, e12)) | (e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e16)) | (e11 = op1(e12, e15)) | (e11 = op1(e12, e14)) | (e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e16, e12)) | (e10 = op1(e15, e12)) | (e10 = op1(e14, e12)) | (e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e16)) | (e10 = op1(e12, e15)) | (e10 = op1(e12, e14)) | (e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e16 = op1(e16, e11)) | (e16 = op1(e15, e11)) | (e16 = op1(e14, e11)) | (e16 = op1(e13, e11)) | (e16 = op1(e12, e11)) | (e16 = op1(e11, e11)) | (e16 = op1(e10, e11))) & ((e16 = op1(e11, e16)) | (e16 = op1(e11, e15)) | (e16 = op1(e11, e14)) | (e16 = op1(e11, e13)) | (e16 = op1(e11, e12)) | (e16 = op1(e11, e11)) | (e16 = op1(e11, e10))) & ((e15 = op1(e16, e11)) | (e15 = op1(e15, e11)) | (e15 = op1(e14, e11)) | (e15 = op1(e13, e11)) | (e15 = op1(e12, e11)) | (e15 = op1(e11, e11)) | (e15 = op1(e10, e11))) & ((e15 = op1(e11, e16)) | (e15 = op1(e11, e15)) | (e15 = op1(e11, e14)) | (e15 = op1(e11, e13)) | (e15 = op1(e11, e12)) | (e15 = op1(e11, e11)) | (e15 = op1(e11, e10))) & ((e14 = op1(e16, e11)) | (e14 = op1(e15, e11)) | (e14 = op1(e14, e11)) | (e14 = op1(e13, e11)) | (e14 = op1(e12, e11)) | (e14 = op1(e11, e11)) | (e14 = op1(e10, e11))) & ((e14 = op1(e11, e16)) | (e14 = op1(e11, e15)) | (e14 = op1(e11, e14)) | (e14 = op1(e11, e13)) | (e14 = op1(e11, e12)) | (e14 = op1(e11, e11)) | (e14 = op1(e11, e10))) & ((e13 = op1(e16, e11)) | (e13 = op1(e15, e11)) | (e13 = op1(e14, e11)) | (e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e16)) | (e13 = op1(e11, e15)) | (e13 = op1(e11, e14)) | (e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e16, e11)) | (e12 = op1(e15, e11)) | (e12 = op1(e14, e11)) | (e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e16)) | (e12 = op1(e11, e15)) | (e12 = op1(e11, e14)) | (e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e16, e11)) | (e11 = op1(e15, e11)) | (e11 = op1(e14, e11)) | (e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e16)) | (e11 = op1(e11, e15)) | (e11 = op1(e11, e14)) | (e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e16, e11)) | (e10 = op1(e15, e11)) | (e10 = op1(e14, e11)) | (e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e16)) | (e10 = op1(e11, e15)) | (e10 = op1(e11, e14)) | (e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e16 = op1(e16, e10)) | (e16 = op1(e15, e10)) | (e16 = op1(e14, e10)) | (e16 = op1(e13, e10)) | (e16 = op1(e12, e10)) | (e16 = op1(e11, e10)) | (op1(e10, e10) = e16)) & ((e16 = op1(e10, e16)) | (e16 = op1(e10, e15)) | (e16 = op1(e10, e14)) | (e16 = op1(e10, e13)) | (e16 = op1(e10, e12)) | (e16 = op1(e10, e11)) | (op1(e10, e10) = e16)) & ((e15 = op1(e16, e10)) | (e15 = op1(e15, e10)) | (e15 = op1(e14, e10)) | (e15 = op1(e13, e10)) | (e15 = op1(e12, e10)) | (e15 = op1(e11, e10)) | (op1(e10, e10) = e15)) & ((e15 = op1(e10, e16)) | (e15 = op1(e10, e15)) | (e15 = op1(e10, e14)) | (e15 = op1(e10, e13)) | (e15 = op1(e10, e12)) | (e15 = op1(e10, e11)) | (op1(e10, e10) = e15)) & ((e14 = op1(e16, e10)) | (e14 = op1(e15, e10)) | (e14 = op1(e14, e10)) | (e14 = op1(e13, e10)) | (e14 = op1(e12, e10)) | (e14 = op1(e11, e10)) | (op1(e10, e10) = e14)) & ((e14 = op1(e10, e16)) | (e14 = op1(e10, e15)) | (e14 = op1(e10, e14)) | (e14 = op1(e10, e13)) | (e14 = op1(e10, e12)) | (e14 = op1(e10, e11)) | (op1(e10, e10) = e14)) & ((e13 = op1(e16, e10)) | (e13 = op1(e15, e10)) | (e13 = op1(e14, e10)) | (e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e16)) | (e13 = op1(e10, e15)) | (e13 = op1(e10, e14)) | (e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e16, e10)) | (e12 = op1(e15, e10)) | (e12 = op1(e14, e10)) | (e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e16)) | (e12 = op1(e10, e15)) | (e12 = op1(e10, e14)) | (e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e16, e10)) | (e11 = op1(e15, e10)) | (e11 = op1(e14, e10)) | (e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e16)) | (e11 = op1(e10, e15)) | (e11 = op1(e10, e14)) | (e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e16, e10)) | (e10 = op1(e15, e10)) | (e10 = op1(e14, e10)) | (e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e16)) | (e10 = op1(e10, e15)) | (e10 = op1(e10, e14)) | (e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax2)).
fof(f3097, plain, (spl42_342 | spl42_293 | spl42_244 | spl42_195 | spl42_146 | spl42_97 | spl42_48), inference(avatar_split_clause, [], [f171, f1785, f1988, f2191, f2394, f2597, f2800, f3003])).
fof(f171, plain, ((e15 = op1(e16, e10)) | (e15 = op1(e15, e10)) | (e15 = op1(e14, e10)) | (e15 = op1(e13, e10)) | (e15 = op1(e12, e10)) | (e15 = op1(e11, e10)) | (op1(e10, e10) = e15)), inference(cnf_transformation, [], [f2])).
fof(f3094, plain, (spl42_288 | spl42_281 | spl42_274 | spl42_267 | spl42_260 | spl42_253 | spl42_246), inference(avatar_split_clause, [], [f174, f2606, f2635, f2664, f2693, f2722, f2751, f2780])).
fof(f174, plain, ((e10 = op1(e11, e16)) | (e10 = op1(e11, e15)) | (e10 = op1(e11, e14)) | (e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f3093, plain, (spl42_330 | spl42_281 | spl42_232 | spl42_183 | spl42_134 | spl42_85 | spl42_36), inference(avatar_split_clause, [], [f175, f1736, f1939, f2142, f2345, f2548, f2751, f2954])).
fof(f175, plain, ((e10 = op1(e16, e11)) | (e10 = op1(e15, e11)) | (e10 = op1(e14, e11)) | (e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f3085, plain, (spl42_334 | spl42_285 | spl42_236 | spl42_187 | spl42_138 | spl42_89 | spl42_40), inference(avatar_split_clause, [], [f183, f1752, f1955, f2158, f2361, f2564, f2767, f2970])).
fof(f183, plain, ((e14 = op1(e16, e11)) | (e14 = op1(e15, e11)) | (e14 = op1(e14, e11)) | (e14 = op1(e13, e11)) | (e14 = op1(e12, e11)) | (e14 = op1(e11, e11)) | (e14 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f3084, plain, (spl42_293 | spl42_286 | spl42_279 | spl42_272 | spl42_265 | spl42_258 | spl42_251), inference(avatar_split_clause, [], [f184, f2626, f2655, f2684, f2713, f2742, f2771, f2800])).
fof(f184, plain, ((e15 = op1(e11, e16)) | (e15 = op1(e11, e15)) | (e15 = op1(e11, e14)) | (e15 = op1(e11, e13)) | (e15 = op1(e11, e12)) | (e15 = op1(e11, e11)) | (e15 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f3081, plain, (spl42_336 | spl42_287 | spl42_238 | spl42_189 | spl42_140 | spl42_91 | spl42_42), inference(avatar_split_clause, [], [f187, f1760, f1963, f2166, f2369, f2572, f2775, f2978])).
fof(f187, plain, ((e16 = op1(e16, e11)) | (e16 = op1(e15, e11)) | (e16 = op1(e14, e11)) | (e16 = op1(e13, e11)) | (e16 = op1(e12, e11)) | (e16 = op1(e11, e11)) | (e16 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f3078, plain, (spl42_240 | spl42_233 | spl42_226 | spl42_219 | spl42_212 | spl42_205 | spl42_198), inference(avatar_split_clause, [], [f190, f2407, f2436, f2465, f2494, f2523, f2552, f2581])).
fof(f190, plain, ((e11 = op1(e12, e16)) | (e11 = op1(e12, e15)) | (e11 = op1(e12, e14)) | (e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f3077, plain, (spl42_324 | spl42_275 | spl42_226 | spl42_177 | spl42_128 | spl42_79 | spl42_30), inference(avatar_split_clause, [], [f191, f1711, f1914, f2117, f2320, f2523, f2726, f2929])).
fof(f191, plain, ((e11 = op1(e16, e12)) | (e11 = op1(e15, e12)) | (e11 = op1(e14, e12)) | (e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f3074, plain, (spl42_242 | spl42_235 | spl42_228 | spl42_221 | spl42_214 | spl42_207 | spl42_200), inference(avatar_split_clause, [], [f194, f2415, f2444, f2473, f2502, f2531, f2560, f2589])).
fof(f194, plain, ((e13 = op1(e12, e16)) | (e13 = op1(e12, e15)) | (e13 = op1(e12, e14)) | (e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f3071, plain, (spl42_327 | spl42_278 | spl42_229 | spl42_180 | spl42_131 | spl42_82 | spl42_33), inference(avatar_split_clause, [], [f197, f1723, f1926, f2129, f2332, f2535, f2738, f2941])).
fof(f197, plain, ((e14 = op1(e16, e12)) | (e14 = op1(e15, e12)) | (e14 = op1(e14, e12)) | (e14 = op1(e13, e12)) | (e14 = op1(e12, e12)) | (e14 = op1(e11, e12)) | (e14 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f3067, plain, (spl42_329 | spl42_280 | spl42_231 | spl42_182 | spl42_133 | spl42_84 | spl42_35), inference(avatar_split_clause, [], [f201, f1731, f1934, f2137, f2340, f2543, f2746, f2949])).
fof(f201, plain, ((e16 = op1(e16, e12)) | (e16 = op1(e15, e12)) | (e16 = op1(e14, e12)) | (e16 = op1(e13, e12)) | (e16 = op1(e12, e12)) | (e16 = op1(e11, e12)) | (e16 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f3066, plain, (spl42_190 | spl42_183 | spl42_176 | spl42_169 | spl42_162 | spl42_155 | spl42_148), inference(avatar_split_clause, [], [f202, f2200, f2229, f2258, f2287, f2316, f2345, f2374])).
fof(f202, plain, ((e10 = op1(e13, e16)) | (e10 = op1(e13, e15)) | (e10 = op1(e13, e14)) | (e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f3064, plain, (spl42_191 | spl42_184 | spl42_177 | spl42_170 | spl42_163 | spl42_156 | spl42_149), inference(avatar_split_clause, [], [f204, f2204, f2233, f2262, f2291, f2320, f2349, f2378])).
fof(f204, plain, ((e11 = op1(e13, e16)) | (e11 = op1(e13, e15)) | (e11 = op1(e13, e14)) | (e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f3053, plain, (spl42_322 | spl42_273 | spl42_224 | spl42_175 | spl42_126 | spl42_77 | spl42_28), inference(avatar_split_clause, [], [f215, f1702, f1905, f2108, f2311, f2514, f2717, f2920])).
fof(f215, plain, ((e16 = op1(e16, e13)) | (e16 = op1(e15, e13)) | (e16 = op1(e14, e13)) | (e16 = op1(e13, e13)) | (e16 = op1(e12, e13)) | (e16 = op1(e11, e13)) | (e16 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f3051, plain, (spl42_309 | spl42_260 | spl42_211 | spl42_162 | spl42_113 | spl42_64 | spl42_15), inference(avatar_split_clause, [], [f217, f1649, f1852, f2055, f2258, f2461, f2664, f2867])).
fof(f217, plain, ((e10 = op1(e16, e14)) | (e10 = op1(e15, e14)) | (e10 = op1(e14, e14)) | (e10 = op1(e13, e14)) | (e10 = op1(e12, e14)) | (e10 = op1(e11, e14)) | (e10 = op1(e10, e14))), inference(cnf_transformation, [], [f2])).
fof(f3045, plain, (spl42_312 | spl42_263 | spl42_214 | spl42_165 | spl42_116 | spl42_67 | spl42_18), inference(avatar_split_clause, [], [f223, f1661, f1864, f2067, f2270, f2473, f2676, f2879])).
fof(f223, plain, ((e13 = op1(e16, e14)) | (e13 = op1(e15, e14)) | (e13 = op1(e14, e14)) | (e13 = op1(e13, e14)) | (e13 = op1(e12, e14)) | (e13 = op1(e11, e14)) | (e13 = op1(e10, e14))), inference(cnf_transformation, [], [f2])).
fof(f3034, plain, (spl42_94 | spl42_87 | spl42_80 | spl42_73 | spl42_66 | spl42_59 | spl42_52), inference(avatar_split_clause, [], [f234, f1802, f1831, f1860, f1889, f1918, f1947, f1976])).
fof(f234, plain, ((e12 = op1(e15, e16)) | (e12 = op1(e15, e15)) | (e12 = op1(e15, e14)) | (e12 = op1(e15, e13)) | (e12 = op1(e15, e12)) | (e12 = op1(e15, e11)) | (e12 = op1(e15, e10))), inference(cnf_transformation, [], [f2])).
fof(f3030, plain, (spl42_96 | spl42_89 | spl42_82 | spl42_75 | spl42_68 | spl42_61 | spl42_54), inference(avatar_split_clause, [], [f238, f1810, f1839, f1868, f1897, f1926, f1955, f1984])).
fof(f238, plain, ((e14 = op1(e15, e16)) | (e14 = op1(e15, e15)) | (e14 = op1(e15, e14)) | (e14 = op1(e15, e13)) | (e14 = op1(e15, e12)) | (e14 = op1(e15, e11)) | (e14 = op1(e15, e10))), inference(cnf_transformation, [], [f2])).
fof(f3028, plain, (spl42_97 | spl42_90 | spl42_83 | spl42_76 | spl42_69 | spl42_62 | spl42_55), inference(avatar_split_clause, [], [f240, f1814, f1843, f1872, f1901, f1930, f1959, f1988])).
fof(f240, plain, ((e15 = op1(e15, e16)) | (e15 = op1(e15, e15)) | (e15 = op1(e15, e14)) | (e15 = op1(e15, e13)) | (e15 = op1(e15, e12)) | (e15 = op1(e15, e11)) | (e15 = op1(e15, e10))), inference(cnf_transformation, [], [f2])).
fof(f3020, plain, (spl42_45 | spl42_38 | spl42_31 | spl42_24 | spl42_17 | spl42_10 | spl42_3), inference(avatar_split_clause, [], [f248, f1599, f1628, f1657, f1686, f1715, f1744, f1773])).
fof(f248, plain, ((e12 = op1(e16, e16)) | (e12 = op1(e16, e15)) | (e12 = op1(e16, e14)) | (e12 = op1(e16, e13)) | (e12 = op1(e16, e12)) | (e12 = op1(e16, e11)) | (e12 = op1(e16, e10))), inference(cnf_transformation, [], [f2])).
fof(f3019, plain, (spl42_297 | spl42_248 | spl42_199 | spl42_150 | spl42_101 | spl42_52 | spl42_3), inference(avatar_split_clause, [], [f249, f1599, f1802, f2005, f2208, f2411, f2614, f2817])).
fof(f249, plain, ((e12 = op1(e16, e16)) | (e12 = op1(e15, e16)) | (e12 = op1(e14, e16)) | (e12 = op1(e13, e16)) | (e12 = op1(e12, e16)) | (e12 = op1(e11, e16)) | (e12 = op1(e10, e16))), inference(cnf_transformation, [], [f2])).
fof(f3012, plain, (spl42_49 | spl42_42 | spl42_35 | spl42_28 | spl42_21 | spl42_14 | spl42_7), inference(avatar_split_clause, [], [f256, f1615, f1644, f1673, f1702, f1731, f1760, f1789])).
fof(f256, plain, ((e16 = op1(e16, e16)) | (e16 = op1(e16, e15)) | (e16 = op1(e16, e14)) | (e16 = op1(e16, e13)) | (e16 = op1(e16, e12)) | (e16 = op1(e16, e11)) | (e16 = op1(e16, e10))), inference(cnf_transformation, [], [f2])).
fof(f2981, plain, (spl42_330 | spl42_331 | spl42_332 | spl42_333 | spl42_334 | spl42_335 | spl42_336), inference(avatar_split_clause, [], [f112, f2978, f2974, f2970, f2966, f2962, f2958, f2954])).
fof(f112, plain, ((e16 = op1(e10, e11)) | (e15 = op1(e10, e11)) | (e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e16 = op1(e16, e16)) | (e15 = op1(e16, e16)) | (e14 = op1(e16, e16)) | (e13 = op1(e16, e16)) | (e12 = op1(e16, e16)) | (e11 = op1(e16, e16)) | (e10 = op1(e16, e16))) & ((e16 = op1(e16, e15)) | (e15 = op1(e16, e15)) | (e14 = op1(e16, e15)) | (e13 = op1(e16, e15)) | (e12 = op1(e16, e15)) | (e11 = op1(e16, e15)) | (e10 = op1(e16, e15))) & ((e16 = op1(e16, e14)) | (e15 = op1(e16, e14)) | (e14 = op1(e16, e14)) | (e13 = op1(e16, e14)) | (e12 = op1(e16, e14)) | (e11 = op1(e16, e14)) | (e10 = op1(e16, e14))) & ((e16 = op1(e16, e13)) | (e15 = op1(e16, e13)) | (e14 = op1(e16, e13)) | (e13 = op1(e16, e13)) | (e12 = op1(e16, e13)) | (e11 = op1(e16, e13)) | (e10 = op1(e16, e13))) & ((e16 = op1(e16, e12)) | (e15 = op1(e16, e12)) | (e14 = op1(e16, e12)) | (e13 = op1(e16, e12)) | (e12 = op1(e16, e12)) | (e11 = op1(e16, e12)) | (e10 = op1(e16, e12))) & ((e16 = op1(e16, e11)) | (e15 = op1(e16, e11)) | (e14 = op1(e16, e11)) | (e13 = op1(e16, e11)) | (e12 = op1(e16, e11)) | (e11 = op1(e16, e11)) | (e10 = op1(e16, e11))) & ((e16 = op1(e16, e10)) | (e15 = op1(e16, e10)) | (e14 = op1(e16, e10)) | (e13 = op1(e16, e10)) | (e12 = op1(e16, e10)) | (e11 = op1(e16, e10)) | (e10 = op1(e16, e10))) & ((e16 = op1(e15, e16)) | (e15 = op1(e15, e16)) | (e14 = op1(e15, e16)) | (e13 = op1(e15, e16)) | (e12 = op1(e15, e16)) | (e11 = op1(e15, e16)) | (e10 = op1(e15, e16))) & ((e16 = op1(e15, e15)) | (e15 = op1(e15, e15)) | (e14 = op1(e15, e15)) | (e13 = op1(e15, e15)) | (e12 = op1(e15, e15)) | (e11 = op1(e15, e15)) | (e10 = op1(e15, e15))) & ((e16 = op1(e15, e14)) | (e15 = op1(e15, e14)) | (e14 = op1(e15, e14)) | (e13 = op1(e15, e14)) | (e12 = op1(e15, e14)) | (e11 = op1(e15, e14)) | (e10 = op1(e15, e14))) & ((e16 = op1(e15, e13)) | (e15 = op1(e15, e13)) | (e14 = op1(e15, e13)) | (e13 = op1(e15, e13)) | (e12 = op1(e15, e13)) | (e11 = op1(e15, e13)) | (e10 = op1(e15, e13))) & ((e16 = op1(e15, e12)) | (e15 = op1(e15, e12)) | (e14 = op1(e15, e12)) | (e13 = op1(e15, e12)) | (e12 = op1(e15, e12)) | (e11 = op1(e15, e12)) | (e10 = op1(e15, e12))) & ((e16 = op1(e15, e11)) | (e15 = op1(e15, e11)) | (e14 = op1(e15, e11)) | (e13 = op1(e15, e11)) | (e12 = op1(e15, e11)) | (e11 = op1(e15, e11)) | (e10 = op1(e15, e11))) & ((e16 = op1(e15, e10)) | (e15 = op1(e15, e10)) | (e14 = op1(e15, e10)) | (e13 = op1(e15, e10)) | (e12 = op1(e15, e10)) | (e11 = op1(e15, e10)) | (e10 = op1(e15, e10))) & ((e16 = op1(e14, e16)) | (e15 = op1(e14, e16)) | (e14 = op1(e14, e16)) | (e13 = op1(e14, e16)) | (e12 = op1(e14, e16)) | (e11 = op1(e14, e16)) | (e10 = op1(e14, e16))) & ((e16 = op1(e14, e15)) | (e15 = op1(e14, e15)) | (e14 = op1(e14, e15)) | (e13 = op1(e14, e15)) | (e12 = op1(e14, e15)) | (e11 = op1(e14, e15)) | (e10 = op1(e14, e15))) & ((e16 = op1(e14, e14)) | (e15 = op1(e14, e14)) | (e14 = op1(e14, e14)) | (e13 = op1(e14, e14)) | (e12 = op1(e14, e14)) | (e11 = op1(e14, e14)) | (e10 = op1(e14, e14))) & ((e16 = op1(e14, e13)) | (e15 = op1(e14, e13)) | (e14 = op1(e14, e13)) | (e13 = op1(e14, e13)) | (e12 = op1(e14, e13)) | (e11 = op1(e14, e13)) | (e10 = op1(e14, e13))) & ((e16 = op1(e14, e12)) | (e15 = op1(e14, e12)) | (e14 = op1(e14, e12)) | (e13 = op1(e14, e12)) | (e12 = op1(e14, e12)) | (e11 = op1(e14, e12)) | (e10 = op1(e14, e12))) & ((e16 = op1(e14, e11)) | (e15 = op1(e14, e11)) | (e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))) & ((e16 = op1(e14, e10)) | (e15 = op1(e14, e10)) | (e14 = op1(e14, e10)) | (e13 = op1(e14, e10)) | (e12 = op1(e14, e10)) | (e11 = op1(e14, e10)) | (e10 = op1(e14, e10))) & ((e16 = op1(e13, e16)) | (e15 = op1(e13, e16)) | (e14 = op1(e13, e16)) | (e13 = op1(e13, e16)) | (e12 = op1(e13, e16)) | (e11 = op1(e13, e16)) | (e10 = op1(e13, e16))) & ((e16 = op1(e13, e15)) | (e15 = op1(e13, e15)) | (e14 = op1(e13, e15)) | (e13 = op1(e13, e15)) | (e12 = op1(e13, e15)) | (e11 = op1(e13, e15)) | (e10 = op1(e13, e15))) & ((e16 = op1(e13, e14)) | (e15 = op1(e13, e14)) | (e14 = op1(e13, e14)) | (e13 = op1(e13, e14)) | (e12 = op1(e13, e14)) | (e11 = op1(e13, e14)) | (e10 = op1(e13, e14))) & ((e16 = op1(e13, e13)) | (e15 = op1(e13, e13)) | (e14 = op1(e13, e13)) | (e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e16 = op1(e13, e12)) | (e15 = op1(e13, e12)) | (e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e16 = op1(e13, e11)) | (e15 = op1(e13, e11)) | (e14 = op1(e13, e11)) | (e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e16 = op1(e13, e10)) | (e15 = op1(e13, e10)) | (e14 = op1(e13, e10)) | (e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e16 = op1(e12, e16)) | (e15 = op1(e12, e16)) | (e14 = op1(e12, e16)) | (e13 = op1(e12, e16)) | (e12 = op1(e12, e16)) | (e11 = op1(e12, e16)) | (e10 = op1(e12, e16))) & ((e16 = op1(e12, e15)) | (e15 = op1(e12, e15)) | (e14 = op1(e12, e15)) | (e13 = op1(e12, e15)) | (e12 = op1(e12, e15)) | (e11 = op1(e12, e15)) | (e10 = op1(e12, e15))) & ((e16 = op1(e12, e14)) | (e15 = op1(e12, e14)) | (e14 = op1(e12, e14)) | (e13 = op1(e12, e14)) | (e12 = op1(e12, e14)) | (e11 = op1(e12, e14)) | (e10 = op1(e12, e14))) & ((e16 = op1(e12, e13)) | (e15 = op1(e12, e13)) | (e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e16 = op1(e12, e12)) | (e15 = op1(e12, e12)) | (e14 = op1(e12, e12)) | (e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e16 = op1(e12, e11)) | (e15 = op1(e12, e11)) | (e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e16 = op1(e12, e10)) | (e15 = op1(e12, e10)) | (e14 = op1(e12, e10)) | (e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e16 = op1(e11, e16)) | (e15 = op1(e11, e16)) | (e14 = op1(e11, e16)) | (e13 = op1(e11, e16)) | (e12 = op1(e11, e16)) | (e11 = op1(e11, e16)) | (e10 = op1(e11, e16))) & ((e16 = op1(e11, e15)) | (e15 = op1(e11, e15)) | (e14 = op1(e11, e15)) | (e13 = op1(e11, e15)) | (e12 = op1(e11, e15)) | (e11 = op1(e11, e15)) | (e10 = op1(e11, e15))) & ((e16 = op1(e11, e14)) | (e15 = op1(e11, e14)) | (e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))) & ((e16 = op1(e11, e13)) | (e15 = op1(e11, e13)) | (e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e16 = op1(e11, e12)) | (e15 = op1(e11, e12)) | (e14 = op1(e11, e12)) | (e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e16 = op1(e11, e11)) | (e15 = op1(e11, e11)) | (e14 = op1(e11, e11)) | (e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e16 = op1(e11, e10)) | (e15 = op1(e11, e10)) | (e14 = op1(e11, e10)) | (e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e16 = op1(e10, e16)) | (e15 = op1(e10, e16)) | (e14 = op1(e10, e16)) | (e13 = op1(e10, e16)) | (e12 = op1(e10, e16)) | (e11 = op1(e10, e16)) | (e10 = op1(e10, e16))) & ((e16 = op1(e10, e15)) | (e15 = op1(e10, e15)) | (e14 = op1(e10, e15)) | (e13 = op1(e10, e15)) | (e12 = op1(e10, e15)) | (e11 = op1(e10, e15)) | (e10 = op1(e10, e15))) & ((e16 = op1(e10, e14)) | (e15 = op1(e10, e14)) | (e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))) & ((e16 = op1(e10, e13)) | (e15 = op1(e10, e13)) | (e14 = op1(e10, e13)) | (e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e16 = op1(e10, e12)) | (e15 = op1(e10, e12)) | (e14 = op1(e10, e12)) | (e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e16 = op1(e10, e11)) | (e15 = op1(e10, e11)) | (e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e16) | (op1(e10, e10) = e15) | (op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG197+1.p', ax1)).
fof(f2923, plain, (spl42_316 | spl42_317 | spl42_318 | spl42_319 | spl42_320 | spl42_321 | spl42_322), inference(avatar_split_clause, [], [f114, f2920, f2916, f2912, f2908, f2904, f2900, f2896])).
fof(f114, plain, ((e16 = op1(e10, e13)) | (e15 = op1(e10, e13)) | (e14 = op1(e10, e13)) | (e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f1])).
fof(f2836, plain, (spl42_295 | spl42_296 | spl42_297 | spl42_298 | spl42_299 | spl42_300 | spl42_301), inference(avatar_split_clause, [], [f117, f2833, f2829, f2825, f2821, f2817, f2813, f2809])).
fof(f117, plain, ((e16 = op1(e10, e16)) | (e15 = op1(e10, e16)) | (e14 = op1(e10, e16)) | (e13 = op1(e10, e16)) | (e12 = op1(e10, e16)) | (e11 = op1(e10, e16)) | (e10 = op1(e10, e16))), inference(cnf_transformation, [], [f1])).
fof(f2720, plain, (spl42_267 | spl42_268 | spl42_269 | spl42_270 | spl42_271 | spl42_272 | spl42_273), inference(avatar_split_clause, [], [f121, f2717, f2713, f2709, f2705, f2701, f2697, f2693])).
fof(f121, plain, ((e16 = op1(e11, e13)) | (e15 = op1(e11, e13)) | (e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))), inference(cnf_transformation, [], [f1])).
fof(f2691, plain, (spl42_260 | spl42_261 | spl42_262 | spl42_263 | spl42_264 | spl42_265 | spl42_266), inference(avatar_split_clause, [], [f122, f2688, f2684, f2680, f2676, f2672, f2668, f2664])).
fof(f122, plain, ((e16 = op1(e11, e14)) | (e15 = op1(e11, e14)) | (e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))), inference(cnf_transformation, [], [f1])).
fof(f2633, plain, (spl42_246 | spl42_247 | spl42_248 | spl42_249 | spl42_250 | spl42_251 | spl42_252), inference(avatar_split_clause, [], [f124, f2630, f2626, f2622, f2618, f2614, f2610, f2606])).
fof(f124, plain, ((e16 = op1(e11, e16)) | (e15 = op1(e11, e16)) | (e14 = op1(e11, e16)) | (e13 = op1(e11, e16)) | (e12 = op1(e11, e16)) | (e11 = op1(e11, e16)) | (e10 = op1(e11, e16))), inference(cnf_transformation, [], [f1])).
fof(f2575, plain, (spl42_232 | spl42_233 | spl42_234 | spl42_235 | spl42_236 | spl42_237 | spl42_238), inference(avatar_split_clause, [], [f126, f2572, f2568, f2564, f2560, f2556, f2552, f2548])).
fof(f126, plain, ((e16 = op1(e12, e11)) | (e15 = op1(e12, e11)) | (e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f2285, plain, (spl42_162 | spl42_163 | spl42_164 | spl42_165 | spl42_166 | spl42_167 | spl42_168), inference(avatar_split_clause, [], [f136, f2282, f2278, f2274, f2270, f2266, f2262, f2258])).
fof(f136, plain, ((e16 = op1(e13, e14)) | (e15 = op1(e13, e14)) | (e14 = op1(e13, e14)) | (e13 = op1(e13, e14)) | (e12 = op1(e13, e14)) | (e11 = op1(e13, e14)) | (e10 = op1(e13, e14))), inference(cnf_transformation, [], [f1])).
fof(f2256, plain, (spl42_155 | spl42_156 | spl42_157 | spl42_158 | spl42_159 | spl42_160 | spl42_161), inference(avatar_split_clause, [], [f137, f2253, f2249, f2245, f2241, f2237, f2233, f2229])).
fof(f137, plain, ((e16 = op1(e13, e15)) | (e15 = op1(e13, e15)) | (e14 = op1(e13, e15)) | (e13 = op1(e13, e15)) | (e12 = op1(e13, e15)) | (e11 = op1(e13, e15)) | (e10 = op1(e13, e15))), inference(cnf_transformation, [], [f1])).
fof(f2227, plain, (spl42_148 | spl42_149 | spl42_150 | spl42_151 | spl42_152 | spl42_153 | spl42_154), inference(avatar_split_clause, [], [f138, f2224, f2220, f2216, f2212, f2208, f2204, f2200])).
fof(f138, plain, ((e16 = op1(e13, e16)) | (e15 = op1(e13, e16)) | (e14 = op1(e13, e16)) | (e13 = op1(e13, e16)) | (e12 = op1(e13, e16)) | (e11 = op1(e13, e16)) | (e10 = op1(e13, e16))), inference(cnf_transformation, [], [f1])).
fof(f2169, plain, (spl42_134 | spl42_135 | spl42_136 | spl42_137 | spl42_138 | spl42_139 | spl42_140), inference(avatar_split_clause, [], [f140, f2166, f2162, f2158, f2154, f2150, f2146, f2142])).
fof(f140, plain, ((e16 = op1(e14, e11)) | (e15 = op1(e14, e11)) | (e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))), inference(cnf_transformation, [], [f1])).
fof(f2053, plain, (spl42_106 | spl42_107 | spl42_108 | spl42_109 | spl42_110 | spl42_111 | spl42_112), inference(avatar_split_clause, [], [f144, f2050, f2046, f2042, f2038, f2034, f2030, f2026])).
fof(f144, plain, ((e16 = op1(e14, e15)) | (e15 = op1(e14, e15)) | (e14 = op1(e14, e15)) | (e13 = op1(e14, e15)) | (e12 = op1(e14, e15)) | (e11 = op1(e14, e15)) | (e10 = op1(e14, e15))), inference(cnf_transformation, [], [f1])).
fof(f2024, plain, (spl42_99 | spl42_100 | spl42_101 | spl42_102 | spl42_103 | spl42_104 | spl42_105), inference(avatar_split_clause, [], [f145, f2021, f2017, f2013, f2009, f2005, f2001, f1997])).
fof(f145, plain, ((e16 = op1(e14, e16)) | (e15 = op1(e14, e16)) | (e14 = op1(e14, e16)) | (e13 = op1(e14, e16)) | (e12 = op1(e14, e16)) | (e11 = op1(e14, e16)) | (e10 = op1(e14, e16))), inference(cnf_transformation, [], [f1])).
fof(f1908, plain, (spl42_71 | spl42_72 | spl42_73 | spl42_74 | spl42_75 | spl42_76 | spl42_77), inference(avatar_split_clause, [], [f149, f1905, f1901, f1897, f1893, f1889, f1885, f1881])).
fof(f149, plain, ((e16 = op1(e15, e13)) | (e15 = op1(e15, e13)) | (e14 = op1(e15, e13)) | (e13 = op1(e15, e13)) | (e12 = op1(e15, e13)) | (e11 = op1(e15, e13)) | (e10 = op1(e15, e13))), inference(cnf_transformation, [], [f1])).
fof(f1879, plain, (spl42_64 | spl42_65 | spl42_66 | spl42_67 | spl42_68 | spl42_69 | spl42_70), inference(avatar_split_clause, [], [f150, f1876, f1872, f1868, f1864, f1860, f1856, f1852])).
fof(f150, plain, ((e16 = op1(e15, e14)) | (e15 = op1(e15, e14)) | (e14 = op1(e15, e14)) | (e13 = op1(e15, e14)) | (e12 = op1(e15, e14)) | (e11 = op1(e15, e14)) | (e10 = op1(e15, e14))), inference(cnf_transformation, [], [f1])).
fof(f1821, plain, (spl42_50 | spl42_51 | spl42_52 | spl42_53 | spl42_54 | spl42_55 | spl42_56), inference(avatar_split_clause, [], [f152, f1818, f1814, f1810, f1806, f1802, f1798, f1794])).
fof(f152, plain, ((e16 = op1(e15, e16)) | (e15 = op1(e15, e16)) | (e14 = op1(e15, e16)) | (e13 = op1(e15, e16)) | (e12 = op1(e15, e16)) | (e11 = op1(e15, e16)) | (e10 = op1(e15, e16))), inference(cnf_transformation, [], [f1])).