fof(f6305, plain, $false, inference(avatar_sat_refutation, [], [f931, f933, f945, f946, f1237, f1238, f1246, f1249, f1251, f1260, f1289, f1313, f1318, f1343, f1352, f1381, f1405, f1415, f1435, f1436, f1437, f1867, f1881, f1895, f3071, f3559, f3567, f3607, f3621, f3634, f3658, f3662, f3670, f3702, f3706, f3725, f3791, f3816, f3838, f3841, f3842, f3886, f3913, f3915, f3922, f3958, f3962, f4014, f4017, f4028, f4030, f4036, f4040, f4051, f4113, f4114, f4148, f4168, f4203, f4210, f4211, f4222, f4224, f4231, f4255, f4263, f4297, f4300, f4351, f4366, f4401, f4453, f4483, f4503, f4517, f4524, f4530, f4605, f4648, f4662, f4707, f4746, f4761, f4814, f4825, f4854, f4883, f4900, f4937, f4976, f4994, f5056, f5102, f5126, f5131, f5140, f5182, f5212, f5271, f5292, f5301, f5310, f5336, f5401, f5483, f5566, f5628, f5681, f5742, f5815, f5884, f5948, f6015, f6073, f6203, f6292])).
fof(f6292, plain, (spl42_122 | ~ spl42_126 | ~ spl42_162), inference(avatar_split_clause, [], [f6288, f1412, f1208, f1191])).
fof(f1191, plain, (spl42_122 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_122])])).
fof(f1208, plain, (spl42_126 <=> (op2(e20, e20) = e21)), introduced(avatar_definition, [new_symbols(naming, [spl42_126])])).
fof(f1412, plain, (spl42_162 <=> (op2(e20, e20) = op2(e20, op2(e20, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl42_162])])).
fof(f6288, plain, ((e21 = op2(e20, e21)) | (~ spl42_126 | ~ spl42_162)), inference(backward_demodulation, [], [f1414, f1210])).
fof(f1210, plain, ((op2(e20, e20) = e21) | ~ spl42_126), inference(avatar_component_clause, [], [f1208])).
fof(f1414, plain, ((op2(e20, e20) = op2(e20, op2(e20, e20))) | ~ spl42_162), inference(avatar_component_clause, [], [f1412])).
fof(f6203, plain, (~ spl42_267 | ~ spl42_46 | spl42_410), inference(avatar_split_clause, [], [f6202, f3052, f836, f2066])).
fof(f2066, plain, (spl42_267 <=> (e21 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_267])])).
fof(f836, plain, (spl42_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_46])])).
fof(f3052, plain, (spl42_410 <=> (h1(e13) = h5(op1(e11, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_410])])).
fof(f6202, plain, (~ (e21 = h1(e13)) | (~ spl42_46 | spl42_410)), inference(forward_demodulation, [], [f6201, f420])).
fof(f420, plain, (e21 = h5(e11)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ((op2(e22, e21) = h5(e13)) & (op2(op2(e22, e21), op2(e22, e21)) = h5(e10)) & (e22 = h5(e12)) & (e21 = h5(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax18)).
fof(f6201, plain, (~ (h1(e13) = h5(e11)) | (~ spl42_46 | spl42_410)), inference(forward_demodulation, [], [f3054, f838])).
fof(f838, plain, ((e11 = op1(e11, e10)) | ~ spl42_46), inference(avatar_component_clause, [], [f836])).
fof(f3054, plain, (~ (h1(e13) = h5(op1(e11, e10))) | spl42_410), inference(avatar_component_clause, [], [f3052])).
fof(f6073, plain, (~ spl42_61 | ~ spl42_125 | spl42_414), inference(avatar_contradiction_clause, [], [f6072])).
fof(f6072, plain, ($false | (~ spl42_61 | ~ spl42_125 | spl42_414)), inference(subsumption_resolution, [], [f6071, f1206])).
fof(f1206, plain, ((e20 = op2(e20, e20)) | ~ spl42_125), inference(avatar_component_clause, [], [f1204])).
fof(f1204, plain, (spl42_125 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_125])])).
fof(f6071, plain, (~ (e20 = op2(e20, e20)) | (~ spl42_61 | spl42_414)), inference(forward_demodulation, [], [f6070, f1474])).
fof(f1474, plain, (e20 = h5(e10)), inference(forward_demodulation, [], [f1473, f1472])).
fof(f1472, plain, (e20 = op2(h5(e13), h5(e13))), inference(backward_demodulation, [], [f402, f423])).
fof(f423, plain, (op2(e22, e21) = h5(e13)), inference(cnf_transformation, [], [f18])).
fof(f402, plain, (e20 = op2(op2(e22, e21), op2(e22, e21))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e23 = op2(e22, e21)) & (e20 = op2(op2(e22, e21), op2(e22, e21)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax13)).
fof(f1473, plain, (h5(e10) = op2(h5(e13), h5(e13))), inference(forward_demodulation, [], [f422, f423])).
fof(f422, plain, (op2(op2(e22, e21), op2(e22, e21)) = h5(e10)), inference(cnf_transformation, [], [f18])).
fof(f6070, plain, (~ (op2(e20, e20) = h5(e10)) | (~ spl42_61 | spl42_414)), inference(forward_demodulation, [], [f3070, f902])).
fof(f902, plain, ((e10 = op1(e10, e10)) | ~ spl42_61), inference(avatar_component_clause, [], [f900])).
fof(f900, plain, (spl42_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_61])])).
fof(f3070, plain, (~ (op2(e20, e20) = h5(op1(e10, e10))) | spl42_414), inference(avatar_component_clause, [], [f3068])).
fof(f3068, plain, (spl42_414 <=> (op2(e20, e20) = h5(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_414])])).
fof(f6015, plain, (~ spl42_58 | ~ spl42_240 | spl42_413), inference(avatar_contradiction_clause, [], [f6014])).
fof(f6014, plain, ($false | (~ spl42_58 | ~ spl42_240 | spl42_413)), inference(subsumption_resolution, [], [f6013, f1918])).
fof(f1918, plain, ((e21 = h4(e13)) | ~ spl42_240), inference(avatar_component_clause, [], [f1917])).
fof(f1917, plain, (spl42_240 <=> (e21 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_240])])).
fof(f6013, plain, (~ (e21 = h4(e13)) | (~ spl42_58 | spl42_413)), inference(forward_demodulation, [], [f6012, f420])).
fof(f6012, plain, (~ (h4(e13) = h5(e11)) | (~ spl42_58 | spl42_413)), inference(forward_demodulation, [], [f3066, f889])).
fof(f889, plain, ((e11 = op1(e10, e11)) | ~ spl42_58), inference(avatar_component_clause, [], [f887])).
fof(f887, plain, (spl42_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_58])])).
fof(f3066, plain, (~ (h4(e13) = h5(op1(e10, e11))) | spl42_413), inference(avatar_component_clause, [], [f3064])).
fof(f3064, plain, (spl42_413 <=> (h4(e13) = h5(op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_413])])).
fof(f5948, plain, (~ spl42_55 | ~ spl42_213 | spl42_412), inference(avatar_contradiction_clause, [], [f5947])).
fof(f5947, plain, ($false | (~ spl42_55 | ~ spl42_213 | spl42_412)), inference(subsumption_resolution, [], [f5946, f421])).
fof(f421, plain, (e22 = h5(e12)), inference(cnf_transformation, [], [f18])).
fof(f5946, plain, (~ (e22 = h5(e12)) | (~ spl42_55 | ~ spl42_213 | spl42_412)), inference(forward_demodulation, [], [f5945, f1764])).
fof(f1764, plain, ((e22 = h7(e13)) | ~ spl42_213), inference(avatar_component_clause, [], [f1763])).
fof(f1763, plain, (spl42_213 <=> (e22 = h7(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_213])])).
fof(f5945, plain, (~ (h5(e12) = h7(e13)) | (~ spl42_55 | spl42_412)), inference(forward_demodulation, [], [f3062, f876])).
fof(f876, plain, ((e12 = op1(e10, e12)) | ~ spl42_55), inference(avatar_component_clause, [], [f874])).
fof(f874, plain, (spl42_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_55])])).
fof(f3062, plain, (~ (h7(e13) = h5(op1(e10, e12))) | spl42_412), inference(avatar_component_clause, [], [f3060])).
fof(f3060, plain, (spl42_412 <=> (h7(e13) = h5(op1(e10, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_412])])).
fof(f5884, plain, (~ spl42_52 | ~ spl42_116 | ~ spl42_398 | spl42_411), inference(avatar_contradiction_clause, [], [f5883])).
fof(f5883, plain, ($false | (~ spl42_52 | ~ spl42_116 | ~ spl42_398 | spl42_411)), inference(subsumption_resolution, [], [f5882, f1167])).
fof(f1167, plain, ((e23 = op2(e20, e23)) | ~ spl42_116), inference(avatar_component_clause, [], [f1165])).
fof(f1165, plain, (spl42_116 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_116])])).
fof(f5882, plain, (~ (e23 = op2(e20, e23)) | (~ spl42_52 | ~ spl42_398 | spl42_411)), inference(forward_demodulation, [], [f5881, f3005])).
fof(f3005, plain, ((e23 = h5(e13)) | ~ spl42_398), inference(avatar_component_clause, [], [f3004])).
fof(f3004, plain, (spl42_398 <=> (e23 = h5(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_398])])).
fof(f5881, plain, (~ (op2(e20, e23) = h5(e13)) | (~ spl42_52 | ~ spl42_398 | spl42_411)), inference(forward_demodulation, [], [f5880, f863])).
fof(f863, plain, ((e13 = op1(e10, e13)) | ~ spl42_52), inference(avatar_component_clause, [], [f861])).
fof(f861, plain, (spl42_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_52])])).
fof(f5880, plain, (~ (op2(e20, e23) = h5(op1(e10, e13))) | (~ spl42_398 | spl42_411)), inference(forward_demodulation, [], [f3058, f3005])).
fof(f3058, plain, (~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | spl42_411), inference(avatar_component_clause, [], [f3056])).
fof(f3056, plain, (spl42_411 <=> (h5(op1(e10, e13)) = op2(e20, h5(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_411])])).
fof(f5815, plain, (~ spl42_40 | ~ spl42_347 | ~ spl42_398 | spl42_408), inference(avatar_contradiction_clause, [], [f5814])).
fof(f5814, plain, ($false | (~ spl42_40 | ~ spl42_347 | ~ spl42_398 | spl42_408)), inference(subsumption_resolution, [], [f5813, f3005])).
fof(f5813, plain, (~ (e23 = h5(e13)) | (~ spl42_40 | ~ spl42_347 | spl42_408)), inference(forward_demodulation, [], [f5812, f2651])).
fof(f2651, plain, ((e23 = h8(e13)) | ~ spl42_347), inference(avatar_component_clause, [], [f2650])).
fof(f2650, plain, (spl42_347 <=> (e23 = h8(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_347])])).
fof(f5812, plain, (~ (h5(e13) = h8(e13)) | (~ spl42_40 | spl42_408)), inference(forward_demodulation, [], [f3046, f812])).
fof(f812, plain, ((e13 = op1(e11, e12)) | ~ spl42_40), inference(avatar_component_clause, [], [f810])).
fof(f810, plain, (spl42_40 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_40])])).
fof(f3046, plain, (~ (h8(e13) = h5(op1(e11, e12))) | spl42_408), inference(avatar_component_clause, [], [f3044])).
fof(f3044, plain, (spl42_408 <=> (h8(e13) = h5(op1(e11, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_408])])).
fof(f5742, plain, (~ spl42_35 | ~ spl42_99 | ~ spl42_398 | spl42_407), inference(avatar_contradiction_clause, [], [f5741])).
fof(f5741, plain, ($false | (~ spl42_35 | ~ spl42_99 | ~ spl42_398 | spl42_407)), inference(subsumption_resolution, [], [f5740, f1095])).
fof(f1095, plain, ((e22 = op2(e21, e23)) | ~ spl42_99), inference(avatar_component_clause, [], [f1093])).
fof(f1093, plain, (spl42_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_99])])).
fof(f5740, plain, (~ (e22 = op2(e21, e23)) | (~ spl42_35 | ~ spl42_398 | spl42_407)), inference(forward_demodulation, [], [f5739, f421])).
fof(f5739, plain, (~ (op2(e21, e23) = h5(e12)) | (~ spl42_35 | ~ spl42_398 | spl42_407)), inference(forward_demodulation, [], [f5738, f791])).
fof(f791, plain, ((e12 = op1(e11, e13)) | ~ spl42_35), inference(avatar_component_clause, [], [f789])).
fof(f789, plain, (spl42_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_35])])).
fof(f5738, plain, (~ (op2(e21, e23) = h5(op1(e11, e13))) | (~ spl42_398 | spl42_407)), inference(forward_demodulation, [], [f3042, f3005])).
fof(f3042, plain, (~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | spl42_407), inference(avatar_component_clause, [], [f3040])).
fof(f3040, plain, (spl42_407 <=> (h5(op1(e11, e13)) = op2(e21, h5(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_407])])).
fof(f5681, plain, (~ spl42_31 | ~ spl42_255 | spl42_406), inference(avatar_contradiction_clause, [], [f5680])).
fof(f5680, plain, ($false | (~ spl42_31 | ~ spl42_255 | spl42_406)), inference(subsumption_resolution, [], [f5679, f2001])).
fof(f2001, plain, ((e22 = h2(e13)) | ~ spl42_255), inference(avatar_component_clause, [], [f2000])).
fof(f2000, plain, (spl42_255 <=> (e22 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_255])])).
fof(f5679, plain, (~ (e22 = h2(e13)) | (~ spl42_31 | spl42_406)), inference(forward_demodulation, [], [f5678, f421])).
fof(f5678, plain, (~ (h2(e13) = h5(e12)) | (~ spl42_31 | spl42_406)), inference(forward_demodulation, [], [f3038, f774])).
fof(f774, plain, ((e12 = op1(e12, e10)) | ~ spl42_31), inference(avatar_component_clause, [], [f772])).
fof(f772, plain, (spl42_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_31])])).
fof(f3038, plain, (~ (h2(e13) = h5(op1(e12, e10))) | spl42_406), inference(avatar_component_clause, [], [f3036])).
fof(f3036, plain, (spl42_406 <=> (h2(e13) = h5(op1(e12, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_406])])).
fof(f5628, plain, (~ spl42_28 | spl42_405), inference(avatar_contradiction_clause, [], [f5627])).
fof(f5627, plain, ($false | (~ spl42_28 | spl42_405)), inference(trivial_inequality_removal, [], [f5626])).
fof(f5626, plain, (~ (h5(e13) = h5(e13)) | (~ spl42_28 | spl42_405)), inference(forward_demodulation, [], [f3034, f761])).
fof(f761, plain, ((e13 = op1(e12, e11)) | ~ spl42_28), inference(avatar_component_clause, [], [f759])).
fof(f759, plain, (spl42_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_28])])).
fof(f3034, plain, (~ (h5(e13) = h5(op1(e12, e11))) | spl42_405), inference(avatar_component_clause, [], [f3032])).
fof(f3032, plain, (spl42_405 <=> (h5(e13) = h5(op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_405])])).
fof(f5566, plain, (~ spl42_21 | ~ spl42_85 | spl42_404), inference(avatar_contradiction_clause, [], [f5565])).
fof(f5565, plain, ($false | (~ spl42_21 | ~ spl42_85 | spl42_404)), inference(subsumption_resolution, [], [f5564, f1036])).
fof(f1036, plain, ((e20 = op2(e22, e22)) | ~ spl42_85), inference(avatar_component_clause, [], [f1034])).
fof(f1034, plain, (spl42_85 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_85])])).
fof(f5564, plain, (~ (e20 = op2(e22, e22)) | (~ spl42_21 | spl42_404)), inference(forward_demodulation, [], [f5563, f1474])).
fof(f5563, plain, (~ (op2(e22, e22) = h5(e10)) | (~ spl42_21 | spl42_404)), inference(forward_demodulation, [], [f3030, f732])).
fof(f732, plain, ((e10 = op1(e12, e12)) | ~ spl42_21), inference(avatar_component_clause, [], [f730])).
fof(f730, plain, (spl42_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_21])])).
fof(f3030, plain, (~ (op2(e22, e22) = h5(op1(e12, e12))) | spl42_404), inference(avatar_component_clause, [], [f3028])).
fof(f3028, plain, (spl42_404 <=> (op2(e22, e22) = h5(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_404])])).
fof(f5483, plain, (~ spl42_18 | ~ spl42_82 | ~ spl42_398 | spl42_403), inference(avatar_contradiction_clause, [], [f5482])).
fof(f5482, plain, ($false | (~ spl42_18 | ~ spl42_82 | ~ spl42_398 | spl42_403)), inference(subsumption_resolution, [], [f5481, f1023])).
fof(f1023, plain, ((e21 = op2(e22, e23)) | ~ spl42_82), inference(avatar_component_clause, [], [f1021])).
fof(f1021, plain, (spl42_82 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_82])])).
fof(f5481, plain, (~ (e21 = op2(e22, e23)) | (~ spl42_18 | ~ spl42_398 | spl42_403)), inference(forward_demodulation, [], [f5480, f420])).
fof(f5480, plain, (~ (op2(e22, e23) = h5(e11)) | (~ spl42_18 | ~ spl42_398 | spl42_403)), inference(forward_demodulation, [], [f5479, f719])).
fof(f719, plain, ((e11 = op1(e12, e13)) | ~ spl42_18), inference(avatar_component_clause, [], [f717])).
fof(f717, plain, (spl42_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_18])])).
fof(f5479, plain, (~ (op2(e22, e23) = h5(op1(e12, e13))) | (~ spl42_398 | spl42_403)), inference(forward_demodulation, [], [f3026, f3005])).
fof(f3026, plain, (~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | spl42_403), inference(avatar_component_clause, [], [f3024])).
fof(f3024, plain, (spl42_403 <=> (h5(op1(e12, e13)) = op2(e22, h5(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_403])])).
fof(f5401, plain, (~ spl42_16 | ~ spl42_80 | ~ spl42_398 | spl42_402), inference(avatar_contradiction_clause, [], [f5400])).
fof(f5400, plain, ($false | (~ spl42_16 | ~ spl42_80 | ~ spl42_398 | spl42_402)), inference(subsumption_resolution, [], [f5399, f1014])).
fof(f1014, plain, ((e23 = op2(e23, e20)) | ~ spl42_80), inference(avatar_component_clause, [], [f1012])).
fof(f1012, plain, (spl42_80 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_80])])).
fof(f5399, plain, (~ (e23 = op2(e23, e20)) | (~ spl42_16 | ~ spl42_398 | spl42_402)), inference(forward_demodulation, [], [f5398, f3005])).
fof(f5398, plain, (~ (op2(e23, e20) = h5(e13)) | (~ spl42_16 | ~ spl42_398 | spl42_402)), inference(forward_demodulation, [], [f5397, f710])).
fof(f710, plain, ((e13 = op1(e13, e10)) | ~ spl42_16), inference(avatar_component_clause, [], [f708])).
fof(f708, plain, (spl42_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_16])])).
fof(f5397, plain, (~ (op2(e23, e20) = h5(op1(e13, e10))) | (~ spl42_398 | spl42_402)), inference(forward_demodulation, [], [f3022, f3005])).
fof(f3022, plain, (~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | spl42_402), inference(avatar_component_clause, [], [f3020])).
fof(f3020, plain, (spl42_402 <=> (h5(op1(e13, e10)) = op2(h5(e13), e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_402])])).
fof(f5336, plain, (~ spl42_6 | spl42_400), inference(avatar_contradiction_clause, [], [f5335])).
fof(f5335, plain, ($false | (~ spl42_6 | spl42_400)), inference(subsumption_resolution, [], [f5334, f420])).
fof(f5334, plain, (~ (e21 = h5(e11)) | (~ spl42_6 | spl42_400)), inference(forward_demodulation, [], [f3014, f668])).
fof(f668, plain, ((e11 = op1(e13, e12)) | ~ spl42_6), inference(avatar_component_clause, [], [f666])).
fof(f666, plain, (spl42_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_6])])).
fof(f3014, plain, (~ (e21 = h5(op1(e13, e12))) | spl42_400), inference(avatar_component_clause, [], [f3012])).
fof(f3012, plain, (spl42_400 <=> (e21 = h5(op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_400])])).
fof(f5310, plain, (~ spl42_41 | ~ spl42_105 | spl42_409), inference(avatar_contradiction_clause, [], [f5309])).
fof(f5309, plain, ($false | (~ spl42_41 | ~ spl42_105 | spl42_409)), inference(subsumption_resolution, [], [f5307, f1474])).
fof(f5307, plain, (~ (e20 = h5(e10)) | (~ spl42_41 | ~ spl42_105 | spl42_409)), inference(backward_demodulation, [], [f5115, f817])).
fof(f817, plain, ((e10 = op1(e11, e11)) | ~ spl42_41), inference(avatar_component_clause, [], [f815])).
fof(f815, plain, (spl42_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_41])])).
fof(f5115, plain, (~ (e20 = h5(op1(e11, e11))) | (~ spl42_105 | spl42_409)), inference(forward_demodulation, [], [f3050, f1121])).
fof(f1121, plain, ((e20 = op2(e21, e21)) | ~ spl42_105), inference(avatar_component_clause, [], [f1119])).
fof(f1119, plain, (spl42_105 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_105])])).
fof(f3050, plain, (~ (op2(e21, e21) = h5(op1(e11, e11))) | spl42_409), inference(avatar_component_clause, [], [f3048])).
fof(f3048, plain, (spl42_409 <=> (op2(e21, e21) = h5(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_409])])).
fof(f5301, plain, (spl42_41 | ~ spl42_46), inference(avatar_split_clause, [], [f5295, f836, f815])).
fof(f5295, plain, ((e10 = op1(e11, e11)) | ~ spl42_46), inference(backward_demodulation, [], [f352, f838])).
fof(f352, plain, (e10 = op1(op1(e11, e10), e11)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ((((op1(e13, e13) = op1(e13, op1(e13, e13))) & (op1(e12, e13) = op1(e13, op1(e12, e13))) & (op1(e11, e13) = op1(e13, op1(e11, e13))) & (op1(e10, e13) = op1(e13, op1(e10, e13)))) | sP2 | sP1 | sP0) & (e13 = op1(op1(e13, e13), e13)) & (e12 = op1(op1(e13, e12), e13)) & (e11 = op1(op1(e13, e11), e13)) & (e10 = op1(op1(e13, e10), e13)) & (e13 = op1(op1(e12, e13), e12)) & (e12 = op1(op1(e12, e12), e12)) & (e11 = op1(op1(e12, e11), e12)) & (e10 = op1(op1(e12, e10), e12)) & (e13 = op1(op1(e11, e13), e11)) & (e12 = op1(op1(e11, e12), e11)) & (e11 = op1(op1(e11, e11), e11)) & (e10 = op1(op1(e11, e10), e11)) & (e13 = op1(op1(e10, e13), e10)) & (e12 = op1(op1(e10, e12), e10)) & (e11 = op1(op1(e10, e11), e10)) & (e10 = op1(op1(e10, e10), e10))), inference(definition_folding, [], [f10, e31, e30, e29])).
fof(f29, plain, (((op1(e13, e10) = op1(e10, op1(e13, e10))) & (op1(e12, e10) = op1(e10, op1(e12, e10))) & (op1(e11, e10) = op1(e10, op1(e11, e10))) & (op1(e10, e10) = op1(e10, op1(e10, e10)))) | ~ sP0), inference(usedef, [], [e29])).
fof(e29, plain, (sP0 <=> ((op1(e13, e10) = op1(e10, op1(e13, e10))) & (op1(e12, e10) = op1(e10, op1(e12, e10))) & (op1(e11, e10) = op1(e10, op1(e11, e10))) & (op1(e10, e10) = op1(e10, op1(e10, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f30, plain, (((op1(e13, e11) = op1(e11, op1(e13, e11))) & (op1(e12, e11) = op1(e11, op1(e12, e11))) & (op1(e11, e11) = op1(e11, op1(e11, e11))) & (op1(e10, e11) = op1(e11, op1(e10, e11)))) | ~ sP1), inference(usedef, [], [e30])).
fof(e30, plain, (sP1 <=> ((op1(e13, e11) = op1(e11, op1(e13, e11))) & (op1(e12, e11) = op1(e11, op1(e12, e11))) & (op1(e11, e11) = op1(e11, op1(e11, e11))) & (op1(e10, e11) = op1(e11, op1(e10, e11))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f31, plain, (((op1(e13, e12) = op1(e12, op1(e13, e12))) & (op1(e12, e12) = op1(e12, op1(e12, e12))) & (op1(e11, e12) = op1(e12, op1(e11, e12))) & (op1(e10, e12) = op1(e12, op1(e10, e12)))) | ~ sP2), inference(usedef, [], [e31])).
fof(e31, plain, (sP2 <=> ((op1(e13, e12) = op1(e12, op1(e13, e12))) & (op1(e12, e12) = op1(e12, op1(e12, e12))) & (op1(e11, e12) = op1(e12, op1(e11, e12))) & (op1(e10, e12) = op1(e12, op1(e10, e12))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f10, plain, ((((op1(e13, e13) = op1(e13, op1(e13, e13))) & (op1(e12, e13) = op1(e13, op1(e12, e13))) & (op1(e11, e13) = op1(e13, op1(e11, e13))) & (op1(e10, e13) = op1(e13, op1(e10, e13)))) | ((op1(e13, e12) = op1(e12, op1(e13, e12))) & (op1(e12, e12) = op1(e12, op1(e12, e12))) & (op1(e11, e12) = op1(e12, op1(e11, e12))) & (op1(e10, e12) = op1(e12, op1(e10, e12)))) | ((op1(e13, e11) = op1(e11, op1(e13, e11))) & (op1(e12, e11) = op1(e11, op1(e12, e11))) & (op1(e11, e11) = op1(e11, op1(e11, e11))) & (op1(e10, e11) = op1(e11, op1(e10, e11)))) | ((op1(e13, e10) = op1(e10, op1(e13, e10))) & (op1(e12, e10) = op1(e10, op1(e12, e10))) & (op1(e11, e10) = op1(e10, op1(e11, e10))) & (op1(e10, e10) = op1(e10, op1(e10, e10))))) & (e13 = op1(op1(e13, e13), e13)) & (e12 = op1(op1(e13, e12), e13)) & (e11 = op1(op1(e13, e11), e13)) & (e10 = op1(op1(e13, e10), e13)) & (e13 = op1(op1(e12, e13), e12)) & (e12 = op1(op1(e12, e12), e12)) & (e11 = op1(op1(e12, e11), e12)) & (e10 = op1(op1(e12, e10), e12)) & (e13 = op1(op1(e11, e13), e11)) & (e12 = op1(op1(e11, e12), e11)) & (e11 = op1(op1(e11, e11), e11)) & (e10 = op1(op1(e11, e10), e11)) & (e13 = op1(op1(e10, e13), e10)) & (e12 = op1(op1(e10, e12), e10)) & (e11 = op1(op1(e10, e11), e10)) & (e10 = op1(op1(e10, e10), e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax10)).
fof(f5292, plain, (spl42_46 | ~ spl42_58), inference(avatar_split_clause, [], [f5287, f887, f836])).
fof(f5287, plain, ((e11 = op1(e11, e10)) | ~ spl42_58), inference(backward_demodulation, [], [f349, f889])).
fof(f349, plain, (e11 = op1(op1(e10, e11), e10)), inference(cnf_transformation, [], [f32])).
fof(f5271, plain, (~ spl42_45 | ~ spl42_61), inference(avatar_split_clause, [], [f5255, f900, f832])).
fof(f832, plain, (spl42_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_45])])).
fof(f5255, plain, (~ (e10 = op1(e11, e10)) | ~ spl42_61), inference(backward_demodulation, [], [f212, f902])).
fof(f212, plain, ~ (op1(e10, e10) = op1(e11, e10)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax5)).
fof(f5212, plain, (~ spl42_45 | spl42_61 | ~ spl42_142), inference(avatar_contradiction_clause, [], [f5211])).
fof(f5211, plain, ($false | (~ spl42_45 | spl42_61 | ~ spl42_142)), inference(subsumption_resolution, [], [f5208, f901])).
fof(f901, plain, (~ (e10 = op1(e10, e10)) | spl42_61), inference(avatar_component_clause, [], [f900])).
fof(f5208, plain, ((e10 = op1(e10, e10)) | (~ spl42_45 | ~ spl42_142)), inference(backward_demodulation, [], [f1317, f834])).
fof(f834, plain, ((e10 = op1(e11, e10)) | ~ spl42_45), inference(avatar_component_clause, [], [f832])).
fof(f1317, plain, ((op1(e11, e10) = op1(e10, op1(e11, e10))) | ~ spl42_142), inference(avatar_component_clause, [], [f1315])).
fof(f1315, plain, (spl42_142 <=> (op1(e11, e10) = op1(e10, op1(e11, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_142])])).
fof(f5182, plain, (spl42_45 | ~ spl42_62), inference(avatar_split_clause, [], [f5167, f904, f832])).
fof(f904, plain, (spl42_62 <=> (op1(e10, e10) = e11)), introduced(avatar_definition, [new_symbols(naming, [spl42_62])])).
fof(f5167, plain, ((e10 = op1(e11, e10)) | ~ spl42_62), inference(backward_demodulation, [], [f348, f906])).
fof(f906, plain, ((op1(e10, e10) = e11) | ~ spl42_62), inference(avatar_component_clause, [], [f904])).
fof(f348, plain, (e10 = op1(op1(e10, e10), e10)), inference(cnf_transformation, [], [f32])).
fof(f5140, plain, (spl42_95 | ~ spl42_213), inference(avatar_split_clause, [], [f5134, f1763, f1076])).
fof(f1076, plain, (spl42_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_95])])).
fof(f5134, plain, ((e22 = op2(e22, e20)) | ~ spl42_213), inference(backward_demodulation, [], [f1485, f1764])).
fof(f1485, plain, (e22 = op2(h7(e13), e20)), inference(backward_demodulation, [], [f382, f431])).
fof(f431, plain, (op2(e20, e22) = h7(e13)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ((op2(e20, e22) = h7(e13)) & (h7(e10) = op2(op2(e20, e22), op2(e20, e22))) & (e20 = h7(e12)) & (e22 = h7(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax20)).
fof(f382, plain, (e22 = op2(op2(e20, e22), e20)), inference(cnf_transformation, [], [f36])).
fof(f36, plain, ((((op2(e23, e23) = op2(e23, op2(e23, e23))) & (op2(e22, e23) = op2(e23, op2(e22, e23))) & (op2(e21, e23) = op2(e23, op2(e21, e23))) & (op2(e20, e23) = op2(e23, op2(e20, e23)))) | sP5 | sP4 | sP3) & (e23 = op2(op2(e23, e23), e23)) & (e22 = op2(op2(e23, e22), e23)) & (e21 = op2(op2(e23, e21), e23)) & (e20 = op2(op2(e23, e20), e23)) & (e23 = op2(op2(e22, e23), e22)) & (e22 = op2(op2(e22, e22), e22)) & (e21 = op2(op2(e22, e21), e22)) & (e20 = op2(op2(e22, e20), e22)) & (e23 = op2(op2(e21, e23), e21)) & (e22 = op2(op2(e21, e22), e21)) & (e21 = op2(op2(e21, e21), e21)) & (e20 = op2(op2(e21, e20), e21)) & (e23 = op2(op2(e20, e23), e20)) & (e22 = op2(op2(e20, e22), e20)) & (e21 = op2(op2(e20, e21), e20)) & (e20 = op2(op2(e20, e20), e20))), inference(definition_folding, [], [f11, e35, e34, e33])).
fof(f33, plain, (((op2(e23, e20) = op2(e20, op2(e23, e20))) & (op2(e22, e20) = op2(e20, op2(e22, e20))) & (op2(e21, e20) = op2(e20, op2(e21, e20))) & (op2(e20, e20) = op2(e20, op2(e20, e20)))) | ~ sP3), inference(usedef, [], [e33])).
fof(e33, plain, (sP3 <=> ((op2(e23, e20) = op2(e20, op2(e23, e20))) & (op2(e22, e20) = op2(e20, op2(e22, e20))) & (op2(e21, e20) = op2(e20, op2(e21, e20))) & (op2(e20, e20) = op2(e20, op2(e20, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f34, plain, (((op2(e23, e21) = op2(e21, op2(e23, e21))) & (op2(e22, e21) = op2(e21, op2(e22, e21))) & (op2(e21, e21) = op2(e21, op2(e21, e21))) & (op2(e20, e21) = op2(e21, op2(e20, e21)))) | ~ sP4), inference(usedef, [], [e34])).
fof(e34, plain, (sP4 <=> ((op2(e23, e21) = op2(e21, op2(e23, e21))) & (op2(e22, e21) = op2(e21, op2(e22, e21))) & (op2(e21, e21) = op2(e21, op2(e21, e21))) & (op2(e20, e21) = op2(e21, op2(e20, e21))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f35, plain, (((op2(e23, e22) = op2(e22, op2(e23, e22))) & (op2(e22, e22) = op2(e22, op2(e22, e22))) & (op2(e21, e22) = op2(e22, op2(e21, e22))) & (op2(e20, e22) = op2(e22, op2(e20, e22)))) | ~ sP5), inference(usedef, [], [e35])).
fof(e35, plain, (sP5 <=> ((op2(e23, e22) = op2(e22, op2(e23, e22))) & (op2(e22, e22) = op2(e22, op2(e22, e22))) & (op2(e21, e22) = op2(e22, op2(e21, e22))) & (op2(e20, e22) = op2(e22, op2(e20, e22))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f11, plain, ((((op2(e23, e23) = op2(e23, op2(e23, e23))) & (op2(e22, e23) = op2(e23, op2(e22, e23))) & (op2(e21, e23) = op2(e23, op2(e21, e23))) & (op2(e20, e23) = op2(e23, op2(e20, e23)))) | ((op2(e23, e22) = op2(e22, op2(e23, e22))) & (op2(e22, e22) = op2(e22, op2(e22, e22))) & (op2(e21, e22) = op2(e22, op2(e21, e22))) & (op2(e20, e22) = op2(e22, op2(e20, e22)))) | ((op2(e23, e21) = op2(e21, op2(e23, e21))) & (op2(e22, e21) = op2(e21, op2(e22, e21))) & (op2(e21, e21) = op2(e21, op2(e21, e21))) & (op2(e20, e21) = op2(e21, op2(e20, e21)))) | ((op2(e23, e20) = op2(e20, op2(e23, e20))) & (op2(e22, e20) = op2(e20, op2(e22, e20))) & (op2(e21, e20) = op2(e20, op2(e21, e20))) & (op2(e20, e20) = op2(e20, op2(e20, e20))))) & (e23 = op2(op2(e23, e23), e23)) & (e22 = op2(op2(e23, e22), e23)) & (e21 = op2(op2(e23, e21), e23)) & (e20 = op2(op2(e23, e20), e23)) & (e23 = op2(op2(e22, e23), e22)) & (e22 = op2(op2(e22, e22), e22)) & (e21 = op2(op2(e22, e21), e22)) & (e20 = op2(op2(e22, e20), e22)) & (e23 = op2(op2(e21, e23), e21)) & (e22 = op2(op2(e21, e22), e21)) & (e21 = op2(op2(e21, e21), e21)) & (e20 = op2(op2(e21, e20), e21)) & (e23 = op2(op2(e20, e23), e20)) & (e22 = op2(op2(e20, e22), e20)) & (e21 = op2(op2(e20, e21), e20)) & (e20 = op2(op2(e20, e20), e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax11)).
fof(f5131, plain, (spl42_119 | ~ spl42_160 | ~ spl42_255), inference(avatar_split_clause, [], [f5122, f2000, f1402, f1178])).
fof(f1178, plain, (spl42_119 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_119])])).
fof(f1402, plain, (spl42_160 <=> (op2(e22, e20) = op2(e20, op2(e22, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl42_160])])).
fof(f5122, plain, ((e22 = op2(e20, e22)) | (~ spl42_160 | ~ spl42_255)), inference(backward_demodulation, [], [f5107, f2001])).
fof(f5107, plain, ((h2(e13) = op2(e20, h2(e13))) | ~ spl42_160), inference(backward_demodulation, [], [f1404, f411])).
fof(f411, plain, (op2(e22, e20) = h2(e13)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((op2(e22, e20) = h2(e13)) & (h2(e10) = op2(op2(e22, e20), op2(e22, e20))) & (e22 = h2(e12)) & (e20 = h2(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax15)).
fof(f1404, plain, ((op2(e22, e20) = op2(e20, op2(e22, e20))) | ~ spl42_160), inference(avatar_component_clause, [], [f1402])).
fof(f5126, plain, (spl42_85 | ~ spl42_255), inference(avatar_split_clause, [], [f5119, f2000, f1034])).
fof(f5119, plain, ((e20 = op2(e22, e22)) | ~ spl42_255), inference(backward_demodulation, [], [f1451, f2001])).
fof(f1451, plain, (e20 = op2(h2(e13), e22)), inference(backward_demodulation, [], [f388, f411])).
fof(f388, plain, (e20 = op2(op2(e22, e20), e22)), inference(cnf_transformation, [], [f36])).
fof(f5102, plain, (spl42_119 | ~ spl42_127 | ~ spl42_162), inference(avatar_split_clause, [], [f5013, f1412, f1212, f1178])).
fof(f1212, plain, (spl42_127 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl42_127])])).
fof(f5013, plain, ((e22 = op2(e20, e22)) | (~ spl42_127 | ~ spl42_162)), inference(forward_demodulation, [], [f1414, f1214])).
fof(f1214, plain, ((op2(e20, e20) = e22) | ~ spl42_127), inference(avatar_component_clause, [], [f1212])).
fof(f5056, plain, (spl42_31 | ~ spl42_55), inference(avatar_split_clause, [], [f5052, f874, f772])).
fof(f5052, plain, ((e12 = op1(e12, e10)) | ~ spl42_55), inference(backward_demodulation, [], [f350, f876])).
fof(f350, plain, (e12 = op1(op1(e10, e12), e10)), inference(cnf_transformation, [], [f32])).
fof(f4994, plain, (~ spl42_1 | spl42_399), inference(avatar_contradiction_clause, [], [f4993])).
fof(f4993, plain, ($false | (~ spl42_1 | spl42_399)), inference(subsumption_resolution, [], [f4992, f1474])).
fof(f4992, plain, (~ (e20 = h5(e10)) | (~ spl42_1 | spl42_399)), inference(forward_demodulation, [], [f3010, f647])).
fof(f647, plain, ((e10 = op1(e13, e13)) | ~ spl42_1), inference(avatar_component_clause, [], [f645])).
fof(f645, plain, (spl42_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_1])])).
fof(f3010, plain, (~ (e20 = h5(op1(e13, e13))) | spl42_399), inference(avatar_component_clause, [], [f3008])).
fof(f3008, plain, (spl42_399 <=> (e20 = h5(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_399])])).
fof(f4976, plain, (~ spl42_126 | ~ spl42_240), inference(avatar_split_clause, [], [f4975, f1917, f1208])).
fof(f4975, plain, (~ (op2(e20, e20) = e21) | ~ spl42_240), inference(forward_demodulation, [], [f1462, f1918])).
fof(f1462, plain, ~ (op2(e20, e20) = h4(e13)), inference(backward_demodulation, [], [f284, f419])).
fof(f419, plain, (op2(e20, e21) = h4(e13)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(e20, e21) = h4(e13)) & (h4(e10) = op2(op2(e20, e21), op2(e20, e21))) & (e20 = h4(e12)) & (e21 = h4(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax17)).
fof(f284, plain, ~ (op2(e20, e20) = op2(e20, e21)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax6)).
fof(f4937, plain, (~ spl42_13 | ~ spl42_1), inference(avatar_split_clause, [], [f4122, f645, f696])).
fof(f696, plain, (spl42_13 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_13])])).
fof(f4122, plain, (~ (e10 = op1(e13, e10)) | ~ spl42_1), inference(forward_demodulation, [], [f256, f647])).
fof(f256, plain, ~ (op1(e13, e10) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4900, plain, (~ spl42_1 | ~ spl42_52 | ~ spl42_147), inference(avatar_contradiction_clause, [], [f4899])).
fof(f4899, plain, ($false | (~ spl42_1 | ~ spl42_52 | ~ spl42_147)), inference(subsumption_resolution, [], [f4898, f310])).
fof(f310, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax7)).
fof(f4898, plain, ((e10 = e13) | (~ spl42_1 | ~ spl42_52 | ~ spl42_147)), inference(forward_demodulation, [], [f4897, f647])).
fof(f4897, plain, ((e13 = op1(e13, e13)) | (~ spl42_52 | ~ spl42_147)), inference(forward_demodulation, [], [f1342, f863])).
fof(f1342, plain, ((op1(e10, e13) = op1(e13, op1(e10, e13))) | ~ spl42_147), inference(avatar_component_clause, [], [f1340])).
fof(f1340, plain, (spl42_147 <=> (op1(e10, e13) = op1(e13, op1(e10, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_147])])).
fof(f4883, plain, (~ spl42_65 | ~ spl42_116 | ~ spl42_166), inference(avatar_contradiction_clause, [], [f4882])).
fof(f4882, plain, ($false | (~ spl42_65 | ~ spl42_116 | ~ spl42_166)), inference(subsumption_resolution, [], [f4881, f316])).
fof(f316, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax8)).
fof(f4881, plain, ((e20 = e23) | (~ spl42_65 | ~ spl42_116 | ~ spl42_166)), inference(forward_demodulation, [], [f4880, f951])).
fof(f951, plain, ((e20 = op2(e23, e23)) | ~ spl42_65), inference(avatar_component_clause, [], [f949])).
fof(f949, plain, (spl42_65 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_65])])).
fof(f4880, plain, ((e23 = op2(e23, e23)) | (~ spl42_116 | ~ spl42_166)), inference(forward_demodulation, [], [f1434, f1167])).
fof(f1434, plain, ((op2(e20, e23) = op2(e23, op2(e20, e23))) | ~ spl42_166), inference(avatar_component_clause, [], [f1432])).
fof(f1432, plain, (spl42_166 <=> (op2(e20, e23) = op2(e23, op2(e20, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl42_166])])).
fof(f4854, plain, (spl42_110 | ~ spl42_240), inference(avatar_split_clause, [], [f4847, f1917, f1140])).
fof(f1140, plain, (spl42_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_110])])).
fof(f4847, plain, ((e21 = op2(e21, e20)) | ~ spl42_240), inference(backward_demodulation, [], [f1465, f1918])).
fof(f1465, plain, (e21 = op2(h4(e13), e20)), inference(backward_demodulation, [], [f381, f419])).
fof(f381, plain, (e21 = op2(op2(e20, e21), e20)), inference(cnf_transformation, [], [f36])).
fof(f4825, plain, (spl42_105 | ~ spl42_267), inference(avatar_split_clause, [], [f4818, f2066, f1119])).
fof(f4818, plain, ((e20 = op2(e21, e21)) | ~ spl42_267), inference(backward_demodulation, [], [f1444, f2067])).
fof(f2067, plain, ((e21 = h1(e13)) | ~ spl42_267), inference(avatar_component_clause, [], [f2066])).
fof(f1444, plain, (e20 = op2(h1(e13), e21)), inference(backward_demodulation, [], [f384, f407])).
fof(f407, plain, (op2(e21, e20) = h1(e13)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((op2(e21, e20) = h1(e13)) & (h1(e10) = op2(op2(e21, e20), op2(e21, e20))) & (e21 = h1(e12)) & (e20 = h1(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax14)).
fof(f384, plain, (e20 = op2(op2(e21, e20), e21)), inference(cnf_transformation, [], [f36])).
fof(f4814, plain, (~ spl42_11 | ~ spl42_75 | ~ spl42_398 | spl42_401), inference(avatar_contradiction_clause, [], [f4813])).
fof(f4813, plain, ($false | (~ spl42_11 | ~ spl42_75 | ~ spl42_398 | spl42_401)), inference(subsumption_resolution, [], [f4812, f993])).
fof(f993, plain, ((e22 = op2(e23, e21)) | ~ spl42_75), inference(avatar_component_clause, [], [f991])).
fof(f991, plain, (spl42_75 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_75])])).
fof(f4812, plain, (~ (e22 = op2(e23, e21)) | (~ spl42_11 | ~ spl42_398 | spl42_401)), inference(forward_demodulation, [], [f4811, f421])).
fof(f4811, plain, (~ (op2(e23, e21) = h5(e12)) | (~ spl42_11 | ~ spl42_398 | spl42_401)), inference(forward_demodulation, [], [f4810, f689])).
fof(f689, plain, ((e12 = op1(e13, e11)) | ~ spl42_11), inference(avatar_component_clause, [], [f687])).
fof(f687, plain, (spl42_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_11])])).
fof(f4810, plain, (~ (op2(e23, e21) = h5(op1(e13, e11))) | (~ spl42_398 | spl42_401)), inference(forward_demodulation, [], [f3018, f3005])).
fof(f3018, plain, (~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | spl42_401), inference(avatar_component_clause, [], [f3016])).
fof(f3016, plain, (spl42_401 <=> (h5(op1(e13, e11)) = op2(h5(e13), e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_401])])).
fof(f4761, plain, (~ spl42_29 | spl42_61 | ~ spl42_141), inference(avatar_contradiction_clause, [], [f4760])).
fof(f4760, plain, ($false | (~ spl42_29 | spl42_61 | ~ spl42_141)), inference(subsumption_resolution, [], [f4759, f901])).
fof(f4759, plain, ((e10 = op1(e10, e10)) | (~ spl42_29 | ~ spl42_141)), inference(forward_demodulation, [], [f1312, f766])).
fof(f766, plain, ((e10 = op1(e12, e10)) | ~ spl42_29), inference(avatar_component_clause, [], [f764])).
fof(f764, plain, (spl42_29 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_29])])).
fof(f1312, plain, ((op1(e12, e10) = op1(e10, op1(e12, e10))) | ~ spl42_141), inference(avatar_component_clause, [], [f1310])).
fof(f1310, plain, (spl42_141 <=> (op1(e12, e10) = op1(e10, op1(e12, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_141])])).
fof(f4746, plain, (~ spl42_92 | spl42_100 | ~ spl42_155), inference(avatar_contradiction_clause, [], [f4745])).
fof(f4745, plain, ($false | (~ spl42_92 | spl42_100 | ~ spl42_155)), inference(subsumption_resolution, [], [f4744, f1098])).
fof(f1098, plain, (~ (e23 = op2(e21, e23)) | spl42_100), inference(avatar_component_clause, [], [f1097])).
fof(f1097, plain, (spl42_100 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_100])])).
fof(f4744, plain, ((e23 = op2(e21, e23)) | (~ spl42_92 | ~ spl42_155)), inference(forward_demodulation, [], [f1380, f1065])).
fof(f1065, plain, ((e23 = op2(e22, e21)) | ~ spl42_92), inference(avatar_component_clause, [], [f1063])).
fof(f1063, plain, (spl42_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_92])])).
fof(f1380, plain, ((op2(e22, e21) = op2(e21, op2(e22, e21))) | ~ spl42_155), inference(avatar_component_clause, [], [f1378])).
fof(f1378, plain, (spl42_155 <=> (op2(e22, e21) = op2(e21, op2(e22, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl42_155])])).
fof(f4707, plain, (spl42_126 | ~ spl42_243), inference(avatar_split_clause, [], [f4701, f1934, f1208])).
fof(f1934, plain, (spl42_243 <=> (e20 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_243])])).
fof(f4701, plain, ((op2(e20, e20) = e21) | ~ spl42_243), inference(backward_demodulation, [], [f1465, f1935])).
fof(f1935, plain, ((e20 = h4(e13)) | ~ spl42_243), inference(avatar_component_clause, [], [f1934])).
fof(f4662, plain, (spl42_99 | ~ spl42_198), inference(avatar_split_clause, [], [f4502, f1681, f1093])).
fof(f1681, plain, (spl42_198 <=> (e21 = h9(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_198])])).
fof(f4502, plain, ((e22 = op2(e21, e23)) | ~ spl42_198), inference(backward_demodulation, [], [f1495, f1682])).
fof(f1682, plain, ((e21 = h9(e13)) | ~ spl42_198), inference(avatar_component_clause, [], [f1681])).
fof(f1495, plain, (e22 = op2(h9(e13), e23)), inference(backward_demodulation, [], [f394, f439])).
fof(f439, plain, (op2(e23, e22) = h9(e13)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ((op2(e23, e22) = h9(e13)) & (h9(e10) = op2(op2(e23, e22), op2(e23, e22))) & (e23 = h9(e12)) & (e22 = h9(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax22)).
fof(f394, plain, (e22 = op2(op2(e23, e22), e23)), inference(cnf_transformation, [], [f36])).
fof(f4648, plain, (spl42_82 | ~ spl42_222), inference(avatar_split_clause, [], [f4489, f1813, f1021])).
fof(f1813, plain, (spl42_222 <=> (e22 = h6(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_222])])).
fof(f4489, plain, ((e21 = op2(e22, e23)) | ~ spl42_222), inference(backward_demodulation, [], [f1478, f1814])).
fof(f1814, plain, ((e22 = h6(e13)) | ~ spl42_222), inference(avatar_component_clause, [], [f1813])).
fof(f1478, plain, (e21 = op2(h6(e13), e23)), inference(backward_demodulation, [], [f393, f427])).
fof(f427, plain, (op2(e23, e21) = h6(e13)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ((op2(e23, e21) = h6(e13)) & (h6(e10) = op2(op2(e23, e21), op2(e23, e21))) & (e23 = h6(e12)) & (e21 = h6(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax19)).
fof(f393, plain, (e21 = op2(op2(e23, e21), e23)), inference(cnf_transformation, [], [f36])).
fof(f4605, plain, (spl42_16 | ~ spl42_52), inference(avatar_split_clause, [], [f4603, f861, f708])).
fof(f4603, plain, ((e13 = op1(e13, e10)) | ~ spl42_52), inference(backward_demodulation, [], [f351, f863])).
fof(f351, plain, (e13 = op1(op1(e10, e13), e10)), inference(cnf_transformation, [], [f32])).
fof(f4530, plain, (spl42_264 | ~ spl42_111), inference(avatar_split_clause, [], [f4529, f1144, f2050])).
fof(f2050, plain, (spl42_264 <=> (e22 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_264])])).
fof(f1144, plain, (spl42_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_111])])).
fof(f4529, plain, ((e22 = h1(e13)) | ~ spl42_111), inference(backward_demodulation, [], [f407, f1146])).
fof(f1146, plain, ((e22 = op2(e21, e20)) | ~ spl42_111), inference(avatar_component_clause, [], [f1144])).
fof(f4524, plain, (~ spl42_123 | spl42_237), inference(avatar_contradiction_clause, [], [f4523])).
fof(f4523, plain, ($false | (~ spl42_123 | spl42_237)), inference(subsumption_resolution, [], [f4522, f1903])).
fof(f1903, plain, (~ (e22 = h4(e13)) | spl42_237), inference(avatar_component_clause, [], [f1901])).
fof(f1901, plain, (spl42_237 <=> (e22 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_237])])).
fof(f4522, plain, ((e22 = h4(e13)) | ~ spl42_123), inference(backward_demodulation, [], [f419, f1197])).
fof(f1197, plain, ((e22 = op2(e20, e21)) | ~ spl42_123), inference(avatar_component_clause, [], [f1195])).
fof(f1195, plain, (spl42_123 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_123])])).
fof(f4517, plain, (~ spl42_28 | ~ spl42_35 | ~ spl42_136), inference(avatar_contradiction_clause, [], [f4516])).
fof(f4516, plain, ($false | (~ spl42_28 | ~ spl42_35 | ~ spl42_136)), inference(subsumption_resolution, [], [f4515, f313])).
fof(f313, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f7])).
fof(f4515, plain, ((e12 = e13) | (~ spl42_28 | ~ spl42_35 | ~ spl42_136)), inference(forward_demodulation, [], [f4514, f791])).
fof(f4514, plain, ((e13 = op1(e11, e13)) | (~ spl42_28 | ~ spl42_136)), inference(forward_demodulation, [], [f1288, f761])).
fof(f1288, plain, ((op1(e12, e11) = op1(e11, op1(e12, e11))) | ~ spl42_136), inference(avatar_component_clause, [], [f1286])).
fof(f1286, plain, (spl42_136 <=> (op1(e12, e11) = op1(e11, op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_136])])).
fof(f4503, plain, (spl42_90 | ~ spl42_149 | ~ spl42_198), inference(avatar_split_clause, [], [f4500, f1681, f1349, f1055])).
fof(f1055, plain, (spl42_90 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_90])])).
fof(f1349, plain, (spl42_149 <=> (op2(e23, e22) = op2(e22, op2(e23, e22)))), introduced(avatar_definition, [new_symbols(naming, [spl42_149])])).
fof(f4500, plain, ((e21 = op2(e22, e21)) | (~ spl42_149 | ~ spl42_198)), inference(backward_demodulation, [], [f3544, f1682])).
fof(f3544, plain, ((h9(e13) = op2(e22, h9(e13))) | ~ spl42_149), inference(forward_demodulation, [], [f1351, f439])).
fof(f1351, plain, ((op2(e23, e22) = op2(e22, op2(e23, e22))) | ~ spl42_149), inference(avatar_component_clause, [], [f1349])).
fof(f4483, plain, (spl42_80 | ~ spl42_311), inference(avatar_split_clause, [], [f4479, f2388, f1012])).
fof(f2388, plain, (spl42_311 <=> (e23 = h10(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_311])])).
fof(f4479, plain, ((e23 = op2(e23, e20)) | ~ spl42_311), inference(backward_demodulation, [], [f1501, f2389])).
fof(f2389, plain, ((e23 = h10(e13)) | ~ spl42_311), inference(avatar_component_clause, [], [f2388])).
fof(f1501, plain, (e23 = op2(h10(e13), e20)), inference(backward_demodulation, [], [f383, f443])).
fof(f443, plain, (op2(e20, e23) = h10(e13)), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ((op2(e20, e23) = h10(e13)) & (h10(e10) = op2(op2(e20, e23), op2(e20, e23))) & (e20 = h10(e12)) & (e23 = h10(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax23)).
fof(f383, plain, (e23 = op2(op2(e20, e23), e20)), inference(cnf_transformation, [], [f36])).
fof(f4453, plain, (spl42_216 | ~ spl42_118), inference(avatar_split_clause, [], [f4452, f1174, f1780])).
fof(f1780, plain, (spl42_216 <=> (e21 = h7(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_216])])).
fof(f1174, plain, (spl42_118 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_118])])).
fof(f4452, plain, ((e21 = h7(e13)) | ~ spl42_118), inference(forward_demodulation, [], [f431, f1176])).
fof(f1176, plain, ((e21 = op2(e20, e22)) | ~ spl42_118), inference(avatar_component_clause, [], [f1174])).
fof(f4401, plain, (spl42_52 | ~ spl42_1), inference(avatar_split_clause, [], [f4110, f645, f861])).
fof(f4110, plain, ((e13 = op1(e10, e13)) | ~ spl42_1), inference(forward_demodulation, [], [f363, f647])).
fof(f363, plain, (e13 = op1(op1(e13, e13), e13)), inference(cnf_transformation, [], [f32])).
fof(f4366, plain, (~ spl42_26 | ~ spl42_28), inference(avatar_contradiction_clause, [], [f4365])).
fof(f4365, plain, ($false | (~ spl42_26 | ~ spl42_28)), inference(subsumption_resolution, [], [f4364, f312])).
fof(f312, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f7])).
fof(f4364, plain, ((e11 = e13) | (~ spl42_26 | ~ spl42_28)), inference(backward_demodulation, [], [f761, f753])).
fof(f753, plain, ((e11 = op1(e12, e11)) | ~ spl42_26), inference(avatar_component_clause, [], [f751])).
fof(f751, plain, (spl42_26 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_26])])).
fof(f4351, plain, (~ spl42_46 | ~ spl42_48), inference(avatar_contradiction_clause, [], [f4350])).
fof(f4350, plain, ($false | (~ spl42_46 | ~ spl42_48)), inference(subsumption_resolution, [], [f4349, f312])).
fof(f4349, plain, ((e11 = e13) | (~ spl42_46 | ~ spl42_48)), inference(forward_demodulation, [], [f846, f838])).
fof(f846, plain, ((e13 = op1(e11, e10)) | ~ spl42_48), inference(avatar_component_clause, [], [f844])).
fof(f844, plain, (spl42_48 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_48])])).
fof(f4300, plain, (~ spl42_216 | ~ spl42_70), inference(avatar_split_clause, [], [f4295, f970, f1780])).
fof(f970, plain, (spl42_70 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_70])])).
fof(f4295, plain, (~ (e21 = h7(e13)) | ~ spl42_70), inference(backward_demodulation, [], [f1482, f972])).
fof(f972, plain, ((e21 = op2(e23, e22)) | ~ spl42_70), inference(avatar_component_clause, [], [f970])).
fof(f1482, plain, ~ (op2(e23, e22) = h7(e13)), inference(backward_demodulation, [], [f274, f431])).
fof(f274, plain, ~ (op2(e20, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f4297, plain, (spl42_198 | ~ spl42_70), inference(avatar_split_clause, [], [f4290, f970, f1681])).
fof(f4290, plain, ((e21 = h9(e13)) | ~ spl42_70), inference(backward_demodulation, [], [f439, f972])).
fof(f4263, plain, (~ spl42_90 | ~ spl42_92), inference(avatar_contradiction_clause, [], [f4262])).
fof(f4262, plain, ($false | (~ spl42_90 | ~ spl42_92)), inference(subsumption_resolution, [], [f4261, f318])).
fof(f318, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f8])).
fof(f4261, plain, ((e21 = e23) | (~ spl42_90 | ~ spl42_92)), inference(backward_demodulation, [], [f1065, f1057])).
fof(f1057, plain, ((e21 = op2(e22, e21)) | ~ spl42_90), inference(avatar_component_clause, [], [f1055])).
fof(f4255, plain, (spl42_255 | ~ spl42_95), inference(avatar_split_clause, [], [f4253, f1076, f2000])).
fof(f4253, plain, ((e22 = h2(e13)) | ~ spl42_95), inference(backward_demodulation, [], [f411, f1078])).
fof(f1078, plain, ((e22 = op2(e22, e20)) | ~ spl42_95), inference(avatar_component_clause, [], [f1076])).
fof(f4231, plain, (spl42_90 | ~ spl42_107), inference(avatar_split_clause, [], [f4225, f1127, f1055])).
fof(f1127, plain, (spl42_107 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_107])])).
fof(f4225, plain, ((e21 = op2(e22, e21)) | ~ spl42_107), inference(backward_demodulation, [], [f385, f1129])).
fof(f1129, plain, ((e22 = op2(e21, e21)) | ~ spl42_107), inference(avatar_component_clause, [], [f1127])).
fof(f385, plain, (e21 = op2(op2(e21, e21), e21)), inference(cnf_transformation, [], [f36])).
fof(f4224, plain, (spl42_472 | ~ spl42_112), inference(avatar_split_clause, [], [f4223, f1148, f3461])).
fof(f3461, plain, (spl42_472 <=> (e23 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_472])])).
fof(f1148, plain, (spl42_112 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_112])])).
fof(f4223, plain, ((e23 = h1(e13)) | ~ spl42_112), inference(backward_demodulation, [], [f407, f1150])).
fof(f1150, plain, ((e23 = op2(e21, e20)) | ~ spl42_112), inference(avatar_component_clause, [], [f1148])).
fof(f4222, plain, (~ spl42_113 | spl42_192), inference(avatar_contradiction_clause, [], [f4221])).
fof(f4221, plain, ($false | (~ spl42_113 | spl42_192)), inference(subsumption_resolution, [], [f4218, f1649])).
fof(f1649, plain, (~ (e20 = h10(e13)) | spl42_192), inference(avatar_component_clause, [], [f1647])).
fof(f1647, plain, (spl42_192 <=> (e20 = h10(e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_192])])).
fof(f4218, plain, ((e20 = h10(e13)) | ~ spl42_113), inference(backward_demodulation, [], [f443, f1155])).
fof(f1155, plain, ((e20 = op2(e20, e23)) | ~ spl42_113), inference(avatar_component_clause, [], [f1153])).
fof(f1153, plain, (spl42_113 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_113])])).
fof(f4211, plain, (~ spl42_237 | ~ spl42_119), inference(avatar_split_clause, [], [f4209, f1178, f1901])).
fof(f4209, plain, (~ (e22 = h4(e13)) | ~ spl42_119), inference(backward_demodulation, [], [f1463, f1180])).
fof(f1180, plain, ((e22 = op2(e20, e22)) | ~ spl42_119), inference(avatar_component_clause, [], [f1178])).
fof(f1463, plain, ~ (op2(e20, e22) = h4(e13)), inference(backward_demodulation, [], [f287, f419])).
fof(f287, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f4210, plain, (spl42_213 | ~ spl42_119), inference(avatar_split_clause, [], [f4208, f1178, f1763])).
fof(f4208, plain, ((e22 = h7(e13)) | ~ spl42_119), inference(backward_demodulation, [], [f431, f1180])).
fof(f4203, plain, (~ spl42_108 | ~ spl42_398), inference(avatar_split_clause, [], [f4202, f3004, f1131])).
fof(f1131, plain, (spl42_108 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_108])])).
fof(f4202, plain, (~ (e23 = op2(e21, e21)) | ~ spl42_398), inference(forward_demodulation, [], [f1467, f3005])).
fof(f1467, plain, ~ (op2(e21, e21) = h5(e13)), inference(backward_demodulation, [], [f269, f423])).
fof(f269, plain, ~ (op2(e21, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f4168, plain, (spl42_65 | ~ spl42_398), inference(avatar_split_clause, [], [f4027, f3004, f949])).
fof(f4027, plain, ((e20 = op2(e23, e23)) | ~ spl42_398), inference(backward_demodulation, [], [f1472, f3005])).
fof(f4148, plain, (~ spl42_44 | ~ spl42_28), inference(avatar_split_clause, [], [f4147, f759, f827])).
fof(f827, plain, (spl42_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_44])])).
fof(f4147, plain, (~ (e13 = op1(e11, e11)) | ~ spl42_28), inference(forward_demodulation, [], [f221, f761])).
fof(f221, plain, ~ (op1(e11, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f4114, plain, (spl42_11 | ~ spl42_40), inference(avatar_split_clause, [], [f4083, f810, f687])).
fof(f4083, plain, ((e12 = op1(e13, e11)) | ~ spl42_40), inference(backward_demodulation, [], [f354, f812])).
fof(f354, plain, (e12 = op1(op1(e11, e12), e11)), inference(cnf_transformation, [], [f32])).
fof(f4113, plain, (~ spl42_1 | ~ spl42_50), inference(avatar_contradiction_clause, [], [f4112])).
fof(f4112, plain, ($false | (~ spl42_1 | ~ spl42_50)), inference(subsumption_resolution, [], [f4111, f312])).
fof(f4111, plain, ((e11 = e13) | (~ spl42_1 | ~ spl42_50)), inference(forward_demodulation, [], [f4110, f855])).
fof(f855, plain, ((e11 = op1(e10, e13)) | ~ spl42_50), inference(avatar_component_clause, [], [f853])).
fof(f853, plain, (spl42_50 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_50])])).
fof(f4051, plain, (~ spl42_91 | ~ spl42_92), inference(avatar_contradiction_clause, [], [f4050])).
fof(f4050, plain, ($false | (~ spl42_91 | ~ spl42_92)), inference(subsumption_resolution, [], [f4049, f319])).
fof(f319, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f4049, plain, ((e22 = e23) | (~ spl42_91 | ~ spl42_92)), inference(backward_demodulation, [], [f1065, f1061])).
fof(f1061, plain, ((e22 = op2(e22, e21)) | ~ spl42_91), inference(avatar_component_clause, [], [f1059])).
fof(f1059, plain, (spl42_91 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_91])])).
fof(f4040, plain, (spl42_347 | ~ spl42_104), inference(avatar_split_clause, [], [f4037, f1114, f2650])).
fof(f1114, plain, (spl42_104 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_104])])).
fof(f4037, plain, ((e23 = h8(e13)) | ~ spl42_104), inference(backward_demodulation, [], [f435, f1116])).
fof(f1116, plain, ((e23 = op2(e21, e22)) | ~ spl42_104), inference(avatar_component_clause, [], [f1114])).
fof(f435, plain, (op2(e21, e22) = h8(e13)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ((op2(e21, e22) = h8(e13)) & (h8(e10) = op2(op2(e21, e22), op2(e21, e22))) & (e21 = h8(e12)) & (e22 = h8(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax21)).
fof(f4036, plain, (spl42_267 | ~ spl42_110), inference(avatar_split_clause, [], [f4035, f1140, f2066])).
fof(f4035, plain, ((e21 = h1(e13)) | ~ spl42_110), inference(backward_demodulation, [], [f407, f1142])).
fof(f1142, plain, ((e21 = op2(e21, e20)) | ~ spl42_110), inference(avatar_component_clause, [], [f1140])).
fof(f4030, plain, (spl42_240 | ~ spl42_122), inference(avatar_split_clause, [], [f4029, f1191, f1917])).
fof(f4029, plain, ((e21 = h4(e13)) | ~ spl42_122), inference(backward_demodulation, [], [f419, f1193])).
fof(f1193, plain, ((e21 = op2(e20, e21)) | ~ spl42_122), inference(avatar_component_clause, [], [f1191])).
fof(f4028, plain, (spl42_70 | ~ spl42_398), inference(avatar_split_clause, [], [f4026, f3004, f970])).
fof(f4026, plain, ((e21 = op2(e23, e22)) | ~ spl42_398), inference(backward_demodulation, [], [f1471, f3005])).
fof(f1471, plain, (e21 = op2(h5(e13), e22)), inference(backward_demodulation, [], [f389, f423])).
fof(f389, plain, (e21 = op2(op2(e22, e21), e22)), inference(cnf_transformation, [], [f36])).
fof(f4017, plain, (spl42_116 | ~ spl42_65), inference(avatar_split_clause, [], [f3693, f949, f1165])).
fof(f3693, plain, ((e23 = op2(e20, e23)) | ~ spl42_65), inference(backward_demodulation, [], [f395, f951])).
fof(f395, plain, (e23 = op2(op2(e23, e23), e23)), inference(cnf_transformation, [], [f36])).
fof(f4014, plain, (~ spl42_264 | ~ spl42_99), inference(avatar_split_clause, [], [f4013, f1093, f2050])).
fof(f4013, plain, (~ (e22 = h1(e13)) | ~ spl42_99), inference(forward_demodulation, [], [f1443, f1095])).
fof(f1443, plain, ~ (op2(e21, e23) = h1(e13)), inference(backward_demodulation, [], [f292, f407])).
fof(f292, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f3962, plain, (~ spl42_54 | ~ spl42_6), inference(avatar_split_clause, [], [f3961, f666, f870])).
fof(f870, plain, (spl42_54 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_54])])).
fof(f3961, plain, (~ (e11 = op1(e10, e12)) | ~ spl42_6), inference(forward_demodulation, [], [f226, f668])).
fof(f226, plain, ~ (op1(e10, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f3958, plain, (spl42_55 | ~ spl42_21), inference(avatar_split_clause, [], [f3845, f730, f874])).
fof(f3845, plain, ((e12 = op1(e10, e12)) | ~ spl42_21), inference(backward_demodulation, [], [f358, f732])).
fof(f358, plain, (e12 = op1(op1(e12, e12), e12)), inference(cnf_transformation, [], [f32])).
fof(f3922, plain, (spl42_18 | ~ spl42_11), inference(avatar_split_clause, [], [f3880, f687, f717])).
fof(f3880, plain, ((e11 = op1(e12, e13)) | ~ spl42_11), inference(backward_demodulation, [], [f361, f689])).
fof(f361, plain, (e11 = op1(op1(e13, e11), e13)), inference(cnf_transformation, [], [f32])).
fof(f3915, plain, (spl42_35 | ~ spl42_6), inference(avatar_split_clause, [], [f3914, f666, f789])).
fof(f3914, plain, ((e12 = op1(e11, e13)) | ~ spl42_6), inference(forward_demodulation, [], [f362, f668])).
fof(f362, plain, (e12 = op1(op1(e13, e12), e13)), inference(cnf_transformation, [], [f32])).
fof(f3913, plain, (spl42_26 | ~ spl42_6 | ~ spl42_130), inference(avatar_split_clause, [], [f3912, f1257, f666, f751])).
fof(f1257, plain, (spl42_130 <=> (op1(e13, e12) = op1(e12, op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_130])])).
fof(f3912, plain, ((e11 = op1(e12, e11)) | (~ spl42_6 | ~ spl42_130)), inference(forward_demodulation, [], [f1259, f668])).
fof(f1259, plain, ((op1(e13, e12) = op1(e12, op1(e13, e12))) | ~ spl42_130), inference(avatar_component_clause, [], [f1257])).
fof(f3886, plain, (~ spl42_11 | ~ spl42_17), inference(avatar_contradiction_clause, [], [f3885])).
fof(f3885, plain, ($false | (~ spl42_11 | ~ spl42_17)), inference(subsumption_resolution, [], [f3884, f308])).
fof(f308, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f7])).
fof(f3884, plain, ((e10 = e11) | (~ spl42_11 | ~ spl42_17)), inference(forward_demodulation, [], [f3880, f715])).
fof(f715, plain, ((e10 = op1(e12, e13)) | ~ spl42_17), inference(avatar_component_clause, [], [f713])).
fof(f713, plain, (spl42_17 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_17])])).
fof(f3842, plain, (spl42_1 | ~ spl42_28), inference(avatar_split_clause, [], [f3837, f759, f645])).
fof(f3837, plain, ((e10 = op1(e13, e13)) | ~ spl42_28), inference(backward_demodulation, [], [f400, f761])).
fof(f400, plain, (e10 = op1(op1(e12, e11), op1(e12, e11))), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e13 = op1(e12, e11)) & (e10 = op1(op1(e12, e11), op1(e12, e11)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax12)).
fof(f3841, plain, (spl42_6 | ~ spl42_28), inference(avatar_split_clause, [], [f3836, f759, f666])).
fof(f3836, plain, ((e11 = op1(e13, e12)) | ~ spl42_28), inference(backward_demodulation, [], [f357, f761])).
fof(f357, plain, (e11 = op1(op1(e12, e11), e12)), inference(cnf_transformation, [], [f32])).
fof(f3838, plain, (~ spl42_12 | ~ spl42_28), inference(avatar_split_clause, [], [f3833, f759, f691])).
fof(f691, plain, (spl42_12 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_12])])).
fof(f3833, plain, (~ (e13 = op1(e13, e11)) | ~ spl42_28), inference(backward_demodulation, [], [f223, f761])).
fof(f223, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f3816, plain, (spl42_12 | ~ spl42_36), inference(avatar_split_clause, [], [f3813, f793, f691])).
fof(f793, plain, (spl42_36 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_36])])).
fof(f3813, plain, ((e13 = op1(e13, e11)) | ~ spl42_36), inference(backward_demodulation, [], [f355, f795])).
fof(f795, plain, ((e13 = op1(e11, e13)) | ~ spl42_36), inference(avatar_component_clause, [], [f793])).
fof(f355, plain, (e13 = op1(op1(e11, e13), e11)), inference(cnf_transformation, [], [f32])).
fof(f3791, plain, (~ spl42_25 | ~ spl42_41), inference(avatar_split_clause, [], [f3786, f815, f747])).
fof(f747, plain, (spl42_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_25])])).
fof(f3786, plain, (~ (e10 = op1(e12, e11)) | ~ spl42_41), inference(backward_demodulation, [], [f221, f817])).
fof(f3725, plain, (~ spl42_29 | ~ spl42_61), inference(avatar_split_clause, [], [f3710, f900, f764])).
fof(f3710, plain, (~ (e10 = op1(e12, e10)) | ~ spl42_61), inference(backward_demodulation, [], [f213, f902])).
fof(f213, plain, ~ (op1(e10, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f3706, plain, (~ spl42_192 | ~ spl42_65), inference(avatar_split_clause, [], [f3697, f949, f1647])).
fof(f3697, plain, (~ (e20 = h10(e13)) | ~ spl42_65), inference(backward_demodulation, [], [f1499, f951])).
fof(f1499, plain, ~ (op2(e23, e23) = h10(e13)), inference(backward_demodulation, [], [f280, f443])).
fof(f280, plain, ~ (op2(e20, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f3702, plain, (~ spl42_65 | ~ spl42_114), inference(avatar_contradiction_clause, [], [f3701])).
fof(f3701, plain, ($false | (~ spl42_65 | ~ spl42_114)), inference(subsumption_resolution, [], [f3700, f318])).
fof(f3700, plain, ((e21 = e23) | (~ spl42_65 | ~ spl42_114)), inference(forward_demodulation, [], [f3693, f1159])).
fof(f1159, plain, ((e21 = op2(e20, e23)) | ~ spl42_114), inference(avatar_component_clause, [], [f1157])).
fof(f1157, plain, (spl42_114 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_114])])).
fof(f3670, plain, (spl42_222 | ~ spl42_75), inference(avatar_split_clause, [], [f3666, f991, f1813])).
fof(f3666, plain, ((e22 = h6(e13)) | ~ spl42_75), inference(backward_demodulation, [], [f427, f993])).
fof(f3662, plain, (~ spl42_79 | ~ spl42_80), inference(avatar_contradiction_clause, [], [f3661])).
fof(f3661, plain, ($false | (~ spl42_79 | ~ spl42_80)), inference(subsumption_resolution, [], [f3660, f319])).
fof(f3660, plain, ((e22 = e23) | (~ spl42_79 | ~ spl42_80)), inference(backward_demodulation, [], [f1014, f1010])).
fof(f1010, plain, ((e22 = op2(e23, e20)) | ~ spl42_79), inference(avatar_component_clause, [], [f1008])).
fof(f1008, plain, (spl42_79 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_79])])).
fof(f3658, plain, (~ spl42_472 | ~ spl42_80), inference(avatar_split_clause, [], [f3655, f1012, f3461])).
fof(f3655, plain, (~ (e23 = h1(e13)) | ~ spl42_80), inference(backward_demodulation, [], [f1440, f1014])).
fof(f1440, plain, ~ (op2(e23, e20) = h1(e13)), inference(backward_demodulation, [], [f264, f407])).
fof(f264, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f3634, plain, (~ spl42_85 | ~ spl42_117), inference(avatar_contradiction_clause, [], [f3633])).
fof(f3633, plain, ($false | (~ spl42_85 | ~ spl42_117)), inference(subsumption_resolution, [], [f3632, f315])).
fof(f315, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f3632, plain, ((e20 = e22) | (~ spl42_85 | ~ spl42_117)), inference(forward_demodulation, [], [f3624, f1172])).
fof(f1172, plain, ((e20 = op2(e20, e22)) | ~ spl42_117), inference(avatar_component_clause, [], [f1170])).
fof(f1170, plain, (spl42_117 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_117])])).
fof(f3624, plain, ((e22 = op2(e20, e22)) | ~ spl42_85), inference(backward_demodulation, [], [f390, f1036])).
fof(f390, plain, (e22 = op2(op2(e22, e22), e22)), inference(cnf_transformation, [], [f36])).
fof(f3621, plain, (spl42_398 | ~ spl42_92), inference(avatar_split_clause, [], [f3618, f1063, f3004])).
fof(f3618, plain, ((e23 = h5(e13)) | ~ spl42_92), inference(backward_demodulation, [], [f423, f1065])).
fof(f3607, plain, (~ spl42_311 | ~ spl42_100), inference(avatar_split_clause, [], [f3603, f1097, f2388])).
fof(f3603, plain, (~ (e23 = h10(e13)) | ~ spl42_100), inference(backward_demodulation, [], [f1497, f1099])).
fof(f1099, plain, ((e23 = op2(e21, e23)) | ~ spl42_100), inference(avatar_component_clause, [], [f1097])).
fof(f1497, plain, ~ (op2(e21, e23) = h10(e13)), inference(backward_demodulation, [], [f278, f443])).
fof(f278, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f3567, plain, (spl42_311 | ~ spl42_116), inference(avatar_split_clause, [], [f3564, f1165, f2388])).
fof(f3564, plain, ((e23 = h10(e13)) | ~ spl42_116), inference(backward_demodulation, [], [f443, f1167])).
fof(f3559, plain, (spl42_243 | ~ spl42_121), inference(avatar_split_clause, [], [f3558, f1187, f1934])).
fof(f1187, plain, (spl42_121 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_121])])).
fof(f3558, plain, ((e20 = h4(e13)) | ~ spl42_121), inference(backward_demodulation, [], [f419, f1189])).
fof(f1189, plain, ((e20 = op2(e20, e21)) | ~ spl42_121), inference(avatar_component_clause, [], [f1187])).
fof(f3071, plain, (spl42_234 | spl42_232 | spl42_230 | ~ spl42_398 | ~ spl42_399 | ~ spl42_400 | ~ spl42_401 | ~ spl42_402 | ~ spl42_403 | ~ spl42_404 | ~ spl42_405 | ~ spl42_406 | ~ spl42_407 | ~ spl42_408 | ~ spl42_409 | ~ spl42_410 | ~ spl42_411 | ~ spl42_412 | ~ spl42_413 | ~ spl42_414), inference(avatar_split_clause, [], [f3002, f3068, f3064, f3060, f3056, f3052, f3048, f3044, f3040, f3036, f3032, f3028, f3024, f3020, f3016, f3012, f3008, f3004, f1858, f1871, f1884])).
fof(f1884, plain, (spl42_234 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl42_234])])).
fof(f1871, plain, (spl42_232 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl42_232])])).
fof(f1858, plain, (spl42_230 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl42_230])])).
fof(f3002, plain, (~ (op2(e20, e20) = h5(op1(e10, e10))) | ~ (h4(e13) = h5(op1(e10, e11))) | ~ (h7(e13) = h5(op1(e10, e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h1(e13) = h5(op1(e11, e10))) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18), inference(forward_demodulation, [], [f3001, f1474])).
fof(f3001, plain, (~ (h4(e13) = h5(op1(e10, e11))) | ~ (h7(e13) = h5(op1(e10, e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h1(e13) = h5(op1(e11, e10))) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f3000, f419])).
fof(f3000, plain, (~ (op2(e20, e21) = h5(op1(e10, e11))) | ~ (h7(e13) = h5(op1(e10, e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h1(e13) = h5(op1(e11, e10))) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2999, f1474])).
fof(f2999, plain, (~ (h5(op1(e10, e11)) = op2(h5(e10), e21)) | ~ (h7(e13) = h5(op1(e10, e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h1(e13) = h5(op1(e11, e10))) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2998, f420])).
fof(f2998, plain, (~ (h7(e13) = h5(op1(e10, e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h1(e13) = h5(op1(e11, e10))) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2997, f431])).
fof(f2997, plain, (~ (op2(e20, e22) = h5(op1(e10, e12))) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h1(e13) = h5(op1(e11, e10))) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2996, f1474])).
fof(f2996, plain, (~ (h5(op1(e10, e12)) = op2(h5(e10), e22)) | ~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h1(e13) = h5(op1(e11, e10))) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2995, f421])).
fof(f2995, plain, (~ (h5(op1(e10, e13)) = op2(e20, h5(e13))) | ~ (h1(e13) = h5(op1(e11, e10))) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2994, f1474])).
fof(f2994, plain, (~ (h1(e13) = h5(op1(e11, e10))) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2993, f407])).
fof(f2993, plain, (~ (op2(e21, e20) = h5(op1(e11, e10))) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2992, f420])).
fof(f2992, plain, (~ (h5(op1(e11, e10)) = op2(h5(e11), e20)) | ~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2991, f1474])).
fof(f2991, plain, (~ (op2(e21, e21) = h5(op1(e11, e11))) | ~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2990, f420])).
fof(f2990, plain, (~ (h8(e13) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2989, f435])).
fof(f2989, plain, (~ (op2(e21, e22) = h5(op1(e11, e12))) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2988, f420])).
fof(f2988, plain, (~ (h5(op1(e11, e12)) = op2(h5(e11), e22)) | ~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2987, f421])).
fof(f2987, plain, (~ (h5(op1(e11, e13)) = op2(e21, h5(e13))) | ~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2986, f420])).
fof(f2986, plain, (~ (h2(e13) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2985, f411])).
fof(f2985, plain, (~ (op2(e22, e20) = h5(op1(e12, e10))) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2984, f421])).
fof(f2984, plain, (~ (h5(op1(e12, e10)) = op2(h5(e12), e20)) | ~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2983, f1474])).
fof(f2983, plain, (~ (h5(e13) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2982, f423])).
fof(f2982, plain, (~ (op2(e22, e21) = h5(op1(e12, e11))) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2981, f421])).
fof(f2981, plain, (~ (h5(op1(e12, e11)) = op2(h5(e12), e21)) | ~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2980, f420])).
fof(f2980, plain, (~ (op2(e22, e22) = h5(op1(e12, e12))) | ~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2979, f421])).
fof(f2979, plain, (~ (h5(op1(e12, e13)) = op2(e22, h5(e13))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2978, f421])).
fof(f2978, plain, (~ (h5(op1(e13, e10)) = op2(h5(e13), e20)) | ~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2977, f1474])).
fof(f2977, plain, (~ (h5(op1(e13, e11)) = op2(h5(e13), e21)) | ~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2976, f420])).
fof(f2976, plain, (~ (e21 = h5(op1(e13, e12))) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2975, f1471])).
fof(f2975, plain, (~ (h5(op1(e13, e12)) = op2(h5(e13), e22)) | ~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f2974, f421])).
fof(f2974, plain, (~ (e20 = h5(op1(e13, e13))) | ~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(forward_demodulation, [], [f615, f1472])).
fof(f615, plain, (~ (e23 = h5(e13)) | sP20 | sP19 | sP18 | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))), inference(cnf_transformation, [], [f73])).
fof(f73, plain, (((~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | sP41 | sP40 | sP39 | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | sP38 | sP37 | sP36 | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | sP35 | sP34 | sP33 | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | sP32 | sP31 | sP30 | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | sP29 | sP28 | sP27 | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | sP26 | sP25 | sP24 | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | sP23 | sP22 | sP21 | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | sP20 | sP19 | sP18 | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP17 | sP16 | sP15 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP14 | sP13 | sP12 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP11 | sP10 | sP9 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP8 | sP7 | sP6 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f28, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37])).
fof(f37, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP6), inference(usedef, [], [e37])).
fof(e37, plain, (sP6 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f38, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP7), inference(usedef, [], [e38])).
fof(e38, plain, (sP7 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f39, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP8), inference(usedef, [], [e39])).
fof(e39, plain, (sP8 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f40, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP9), inference(usedef, [], [e40])).
fof(e40, plain, (sP9 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f41, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP10), inference(usedef, [], [e41])).
fof(e41, plain, (sP10 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f42, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP11), inference(usedef, [], [e42])).
fof(e42, plain, (sP11 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f43, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP12), inference(usedef, [], [e43])).
fof(e43, plain, (sP12 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f44, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP13), inference(usedef, [], [e44])).
fof(e44, plain, (sP13 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f45, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP14), inference(usedef, [], [e45])).
fof(e45, plain, (sP14 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f46, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP15), inference(usedef, [], [e46])).
fof(e46, plain, (sP15 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f47, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP16), inference(usedef, [], [e47])).
fof(e47, plain, (sP16 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f48, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP17), inference(usedef, [], [e48])).
fof(e48, plain, (sP17 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f49, plain, ((~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP18), inference(usedef, [], [e49])).
fof(e49, plain, (sP18 <=> (~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f50, plain, ((~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP19), inference(usedef, [], [e50])).
fof(e50, plain, (sP19 <=> (~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f51, plain, ((~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP20), inference(usedef, [], [e51])).
fof(e51, plain, (sP20 <=> (~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f52, plain, ((~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP21), inference(usedef, [], [e52])).
fof(e52, plain, (sP21 <=> (~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f53, plain, ((~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP22), inference(usedef, [], [e53])).
fof(e53, plain, (sP22 <=> (~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f54, plain, ((~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP23), inference(usedef, [], [e54])).
fof(e54, plain, (sP23 <=> (~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f55, plain, ((~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP24), inference(usedef, [], [e55])).
fof(e55, plain, (sP24 <=> (~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f56, plain, ((~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP25), inference(usedef, [], [e56])).
fof(e56, plain, (sP25 <=> (~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f57, plain, ((~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP26), inference(usedef, [], [e57])).
fof(e57, plain, (sP26 <=> (~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f58, plain, ((~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ sP27), inference(usedef, [], [e58])).
fof(e58, plain, (sP27 <=> (~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f59, plain, ((~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | ~ sP28), inference(usedef, [], [e59])).
fof(e59, plain, (sP28 <=> (~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f60, plain, ((~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | ~ sP29), inference(usedef, [], [e60])).
fof(e60, plain, (sP29 <=> (~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f61, plain, ((~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ sP30), inference(usedef, [], [e61])).
fof(e61, plain, (sP30 <=> (~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f62, plain, ((~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | ~ sP31), inference(usedef, [], [e62])).
fof(e62, plain, (sP31 <=> (~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f63, plain, ((~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | ~ sP32), inference(usedef, [], [e63])).
fof(e63, plain, (sP32 <=> (~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f64, plain, ((~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ sP33), inference(usedef, [], [e64])).
fof(e64, plain, (sP33 <=> (~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f65, plain, ((~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | ~ sP34), inference(usedef, [], [e65])).
fof(e65, plain, (sP34 <=> (~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f66, plain, ((~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | ~ sP35), inference(usedef, [], [e66])).
fof(e66, plain, (sP35 <=> (~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f67, plain, ((~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ sP36), inference(usedef, [], [e67])).
fof(e67, plain, (sP36 <=> (~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f68, plain, ((~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | ~ sP37), inference(usedef, [], [e68])).
fof(e68, plain, (sP37 <=> (~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f69, plain, ((~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | ~ sP38), inference(usedef, [], [e69])).
fof(e69, plain, (sP38 <=> (~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f70, plain, ((~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ sP39), inference(usedef, [], [e70])).
fof(e70, plain, (sP39 <=> (~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f71, plain, ((~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | ~ sP40), inference(usedef, [], [e71])).
fof(e71, plain, (sP40 <=> (~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f72, plain, ((~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | ~ sP41), inference(usedef, [], [e72])).
fof(e72, plain, (sP41 <=> (~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f28, plain, (((~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | (~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | (~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | (~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | (~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | (~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | (~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | (~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | (~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | (~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | (~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | (~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | (~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | (~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | (~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | (~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | (~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | (~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | (~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((((e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f26])).
fof(f26, plain, ~ ((((e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', co1)).
fof(f1895, plain, ~ spl42_234, inference(avatar_split_clause, [], [f1894, f1884])).
fof(f1894, plain, ~ sP18, inference(subsumption_resolution, [], [f544, f1474])).
fof(f544, plain, (~ (e20 = h5(e10)) | ~ sP18), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP18), inference(nnf_transformation, [], [f49])).
fof(f1881, plain, ~ spl42_232, inference(avatar_split_clause, [], [f1880, f1871])).
fof(f1880, plain, ~ sP19, inference(subsumption_resolution, [], [f541, f420])).
fof(f541, plain, (~ (e21 = h5(e11)) | ~ sP19), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP19), inference(nnf_transformation, [], [f50])).
fof(f1867, plain, ~ spl42_230, inference(avatar_split_clause, [], [f1866, f1858])).
fof(f1866, plain, ~ sP20, inference(subsumption_resolution, [], [f538, f421])).
fof(f538, plain, (~ (e22 = h5(e12)) | ~ sP20), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP20), inference(nnf_transformation, [], [f51])).
fof(f1437, plain, spl42_92, inference(avatar_split_clause, [], [f403, f1063])).
fof(f403, plain, (e23 = op2(e22, e21)), inference(cnf_transformation, [], [f13])).
fof(f1436, plain, spl42_28, inference(avatar_split_clause, [], [f401, f759])).
fof(f401, plain, (e13 = op1(e12, e11)), inference(cnf_transformation, [], [f12])).
fof(f1435, plain, (spl42_158 | spl42_153 | spl42_148 | spl42_166), inference(avatar_split_clause, [], [f396, f1432, f1345, f1369, f1393])).
fof(f1393, plain, (spl42_158 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl42_158])])).
fof(f1369, plain, (spl42_153 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl42_153])])).
fof(f1345, plain, (spl42_148 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl42_148])])).
fof(f396, plain, ((op2(e20, e23) = op2(e23, op2(e20, e23))) | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f36])).
fof(f1415, plain, (~ spl42_158 | spl42_162), inference(avatar_split_clause, [], [f376, f1412, f1393])).
fof(f376, plain, ((op2(e20, e20) = op2(e20, op2(e20, e20))) | ~ sP3), inference(cnf_transformation, [], [f79])).
fof(f79, plain, (((op2(e23, e20) = op2(e20, op2(e23, e20))) & (op2(e22, e20) = op2(e20, op2(e22, e20))) & (op2(e21, e20) = op2(e20, op2(e21, e20))) & (op2(e20, e20) = op2(e20, op2(e20, e20)))) | ~ sP3), inference(nnf_transformation, [], [f33])).
fof(f1405, plain, (~ spl42_158 | spl42_160), inference(avatar_split_clause, [], [f378, f1402, f1393])).
fof(f378, plain, ((op2(e22, e20) = op2(e20, op2(e22, e20))) | ~ sP3), inference(cnf_transformation, [], [f79])).
fof(f1381, plain, (~ spl42_153 | spl42_155), inference(avatar_split_clause, [], [f374, f1378, f1369])).
fof(f374, plain, ((op2(e22, e21) = op2(e21, op2(e22, e21))) | ~ sP4), inference(cnf_transformation, [], [f78])).
fof(f78, plain, (((op2(e23, e21) = op2(e21, op2(e23, e21))) & (op2(e22, e21) = op2(e21, op2(e22, e21))) & (op2(e21, e21) = op2(e21, op2(e21, e21))) & (op2(e20, e21) = op2(e21, op2(e20, e21)))) | ~ sP4), inference(nnf_transformation, [], [f34])).
fof(f1352, plain, (~ spl42_148 | spl42_149), inference(avatar_split_clause, [], [f371, f1349, f1345])).
fof(f371, plain, ((op2(e23, e22) = op2(e22, op2(e23, e22))) | ~ sP5), inference(cnf_transformation, [], [f77])).
fof(f77, plain, (((op2(e23, e22) = op2(e22, op2(e23, e22))) & (op2(e22, e22) = op2(e22, op2(e22, e22))) & (op2(e21, e22) = op2(e22, op2(e21, e22))) & (op2(e20, e22) = op2(e22, op2(e20, e22)))) | ~ sP5), inference(nnf_transformation, [], [f35])).
fof(f1343, plain, (spl42_139 | spl42_134 | spl42_129 | spl42_147), inference(avatar_split_clause, [], [f364, f1340, f1253, f1277, f1301])).
fof(f1301, plain, (spl42_139 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl42_139])])).
fof(f1277, plain, (spl42_134 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl42_134])])).
fof(f1253, plain, (spl42_129 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl42_129])])).
fof(f364, plain, ((op1(e10, e13) = op1(e13, op1(e10, e13))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f32])).
fof(f1318, plain, (~ spl42_139 | spl42_142), inference(avatar_split_clause, [], [f345, f1315, f1301])).
fof(f345, plain, ((op1(e11, e10) = op1(e10, op1(e11, e10))) | ~ sP0), inference(cnf_transformation, [], [f76])).
fof(f76, plain, (((op1(e13, e10) = op1(e10, op1(e13, e10))) & (op1(e12, e10) = op1(e10, op1(e12, e10))) & (op1(e11, e10) = op1(e10, op1(e11, e10))) & (op1(e10, e10) = op1(e10, op1(e10, e10)))) | ~ sP0), inference(nnf_transformation, [], [f29])).
fof(f1313, plain, (~ spl42_139 | spl42_141), inference(avatar_split_clause, [], [f346, f1310, f1301])).
fof(f346, plain, ((op1(e12, e10) = op1(e10, op1(e12, e10))) | ~ sP0), inference(cnf_transformation, [], [f76])).
fof(f1289, plain, (~ spl42_134 | spl42_136), inference(avatar_split_clause, [], [f342, f1286, f1277])).
fof(f342, plain, ((op1(e12, e11) = op1(e11, op1(e12, e11))) | ~ sP1), inference(cnf_transformation, [], [f75])).
fof(f75, plain, (((op1(e13, e11) = op1(e11, op1(e13, e11))) & (op1(e12, e11) = op1(e11, op1(e12, e11))) & (op1(e11, e11) = op1(e11, op1(e11, e11))) & (op1(e10, e11) = op1(e11, op1(e10, e11)))) | ~ sP1), inference(nnf_transformation, [], [f30])).
fof(f1260, plain, (~ spl42_129 | spl42_130), inference(avatar_split_clause, [], [f339, f1257, f1253])).
fof(f339, plain, ((op1(e13, e12) = op1(e12, op1(e13, e12))) | ~ sP2), inference(cnf_transformation, [], [f74])).
fof(f74, plain, (((op1(e13, e12) = op1(e12, op1(e13, e12))) & (op1(e12, e12) = op1(e12, op1(e12, e12))) & (op1(e11, e12) = op1(e12, op1(e11, e12))) & (op1(e10, e12) = op1(e12, op1(e10, e12)))) | ~ sP2), inference(nnf_transformation, [], [f31])).
fof(f1251, plain, (spl42_125 | spl42_121 | spl42_117 | spl42_113), inference(avatar_split_clause, [], [f180, f1153, f1170, f1187, f1204])).
fof(f180, plain, ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax4)).
fof(f1249, plain, (spl42_126 | spl42_122 | spl42_118 | spl42_114), inference(avatar_split_clause, [], [f182, f1157, f1174, f1191, f1208])).
fof(f182, plain, ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)), inference(cnf_transformation, [], [f4])).
fof(f1246, plain, (spl42_127 | spl42_111 | spl42_95 | spl42_79), inference(avatar_split_clause, [], [f185, f1008, f1076, f1144, f1212])).
fof(f185, plain, ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)), inference(cnf_transformation, [], [f4])).
fof(f1238, plain, (spl42_123 | spl42_107 | spl42_91 | spl42_75), inference(avatar_split_clause, [], [f193, f991, f1059, f1127, f1195])).
fof(f193, plain, ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1237, plain, (spl42_112 | spl42_108 | spl42_104 | spl42_100), inference(avatar_split_clause, [], [f194, f1097, f1114, f1131, f1148])).
fof(f194, plain, ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f946, plain, (spl42_61 | spl42_45 | spl42_29 | spl42_13), inference(avatar_split_clause, [], [f133, f696, f764, f832, f900])).
fof(f133, plain, ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG105+1.p', ax2)).
fof(f945, plain, (spl42_62 | spl42_58 | spl42_54 | spl42_50), inference(avatar_split_clause, [], [f134, f853, f870, f887, f904])).
fof(f134, plain, ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)), inference(cnf_transformation, [], [f2])).
fof(f933, plain, (spl42_48 | spl42_44 | spl42_40 | spl42_36), inference(avatar_split_clause, [], [f146, f793, f810, f827, f844])).
fof(f146, plain, ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f931, plain, (spl42_29 | spl42_25 | spl42_21 | spl42_17), inference(avatar_split_clause, [], [f148, f713, f730, f747, f764])).
fof(f148, plain, ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).