fof(f4854, plain, $false, inference(avatar_sat_refutation, [], [f617, f634, f651, f668, f685, f702, f719, f736, f753, f770, f787, f804, f821, f838, f855, f872, f873, f874, f875, f876, f877, f878, f879, f880, f881, f882, f883, f884, f885, f886, f887, f889, f890, f891, f894, f895, f896, f897, f898, f899, f900, f901, f902, f904, f909, f910, f911, f916, f921, f926, f931, f936, f937, f938, f943, f948, f953, f958, f963, f964, f965, f970, f975, f980, f985, f992, f993, f1000, f1001, f1008, f1009, f1015, f1016, f1022, f1023, f1025, f1030, f1031, f1032, f1033, f1038, f1039, f1040, f1041, f1047, f1048, f1049, f1054, f1055, f1056, f1057, f1062, f1063, f1070, f1071, f1072, f1073, f1079, f1080, f1086, f1087, f1088, f1089, f1094, f1095, f1096, f1097, f1102, f1103, f1110, f1111, f1119, f1120, f1126, f1127, f1128, f1129, f1134, f1135, f1136, f1137, f1144, f1145, f1152, f1153, f1160, f1161, f1168, f1169, f1174, f1175, f1176, f1177, f1182, f1183, f1184, f1185, f1190, f1191, f1198, f1199, f1200, f1201, f1206, f1207, f1208, f1209, f1215, f1216, f1217, f1222, f1223, f1224, f1225, f1230, f1231, f1238, f1239, f1240, f1246, f1247, f1248, f1249, f1255, f1256, f1262, f1263, f1264, f1265, f1270, f1271, f1272, f1273, f1278, f1279, f1287, f1288, f1289, f1294, f1295, f1296, f1297, f1304, f1305, f1312, f1313, f1320, f1321, f1328, f1329, f1334, f1335, f1336, f1337, f1342, f1343, f1344, f1345, f1351, f1352, f1353, f1358, f1359, f1361, f1366, f1367, f1374, f1375, f1376, f1377, f1382, f1383, f1384, f1385, f1391, f1392, f1398, f1399, f1400, f1401, f1406, f1407, f1409, f1414, f1415, f1416, f1417, f1423, f1424, f1425, f1430, f1431, f1432, f1433, f1438, f1439, f1440, f1441, f1446, f1447, f1455, f1456, f1464, f1465, f1472, f1473, f1479, f1480, f1481, f1488, f1489, f1490, f1491, f1492, f1497, f1502, f1507, f1512, f1515, f1516, f1525, f1530, f1539, f1544, f1553, f1558, f1567, f1572, f1577, f1582, f1587, f1592, f1597, f1602, f1607, f1612, f1617, f1618, f1623, f1624, f1629, f1630, f1635, f1636, f1646, f1648, f1651, f1654, f1660, f1669, f1671, f1672, f1699, f1701, f1707, f1711, f1713, f1717, f1727, f1734, f1735, f1751, f1754, f1789, f1799, f1810, f1818, f1834, f1856, f1860, f1885, f1888, f1915, f1923, f1927, f1932, f1944, f1958, f1986, f1989, f2039, f2056, f2063, f2067, f2070, f2074, f2124, f2163, f2171, f2174, f2183, f2221, f2222, f2235, f2258, f2270, f2280, f2292, f2300, f2303, f2312, f2320, f2328, f2333, f2341, f2355, f2357, f2376, f2379, f2393, f2415, f2447, f2467, f2485, f2491, f2505, f2507, f2522, f2528, f2561, f2563, f2566, f2576, f2585, f2586, f2589, f2623, f2627, f2647, f2674, f2686, f2698, f2704, f2718, f2739, f2743, f2746, f2765, f2796, f2814, f2834, f2863, f2888, f2890, f2905, f2913, f2936, f2943, f2945, f2958, f2959, f2965, f2970, f2976, f2993, f2994, f3003, f3006, f3016, f3033, f3043, f3049, f3085, f3102, f3106, f3112, f3119, f3137, f3158, f3166, f3184, f3192, f3206, f3208, f3215, f3234, f3255, f3265, f3301, f3307, f3347, f3354, f3374, f3390, f3396, f3407, f3408, f3449, f3458, f3466, f3478, f3487, f3493, f3500, f3501, f3504, f3510, f3532, f3542, f3557, f3571, f3577, f3587, f3602, f3613, f3625, f3627, f3665, f3696, f3701, f3720, f3728, f3733, f3748, f3752, f3759, f3766, f3768, f3793, f3795, f3804, f3824, f3833, f3840, f3845, f3851, f3884, f3905, f3913, f3923, f3933, f3940, f3955, f3975, f3985, f4010, f4017, f4023, f4045, f4046, f4060, f4092, f4096, f4101, f4109, f4116, f4131, f4139, f4145, f4150, f4152, f4153, f4160, f4169, f4180, f4186, f4188, f4198, f4205, f4218, f4234, f4250, f4251, f4259, f4260, f4316, f4317, f4340, f4355, f4359, f4361, f4367, f4376, f4378, f4399, f4400, f4402, f4411, f4418, f4428, f4434, f4435, f4446, f4451, f4457, f4463, f4464, f4471, f4478, f4492, f4494, f4497, f4499, f4511, f4516, f4526, f4537, f4538, f4552, f4561, f4562, f4567, f4575, f4583, f4600, f4608, f4613, f4621, f4623, f4631, f4636, f4643, f4656, f4657, f4669, f4670, f4678, f4682, f4689, f4690, f4697, f4709, f4719, f4724, f4725, f4732, f4736, f4738, f4744, f4748, f4750, f4752, f4765, f4767, f4773, f4774, f4781, f4791, f4802, f4814, f4816, f4822, f4823, f4831, f4834, f4844])).
fof(f4844, plain, (~ spl66_2 | ~ spl66_10), inference(avatar_split_clause, [], [f4838, f640, f606])).
fof(f606, plain, (spl66_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_2])])).
fof(f640, plain, (spl66_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_10])])).
fof(f4838, plain, (~ (e1 = op(e3, e3)) | ~ spl66_10), inference(backward_demodulation, [], [f281, f642])).
fof(f642, plain, ((e1 = op(e3, e1)) | ~ spl66_10), inference(avatar_component_clause, [], [f640])).
fof(f281, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax3)).
fof(f4834, plain, (~ spl66_21 | spl66_53 | ~ spl66_160), inference(avatar_contradiction_clause, [], [f4833])).
fof(f4833, plain, ($false | (~ spl66_21 | spl66_53 | ~ spl66_160)), inference(subsumption_resolution, [], [f4830, f824])).
fof(f824, plain, (~ (e0 = op(e0, e2)) | spl66_53), inference(avatar_component_clause, [], [f823])).
fof(f823, plain, (spl66_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_53])])).
fof(f4830, plain, ((e0 = op(e0, e2)) | (~ spl66_21 | ~ spl66_160)), inference(backward_demodulation, [], [f1580, f689])).
fof(f689, plain, ((e0 = op(e2, e2)) | ~ spl66_21), inference(avatar_component_clause, [], [f687])).
fof(f687, plain, (spl66_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_21])])).
fof(f1580, plain, ((e0 = op(op(e2, e2), e2)) | ~ spl66_160), inference(avatar_component_clause, [], [f1579])).
fof(f1579, plain, (spl66_160 <=> (e0 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_160])])).
fof(f4831, plain, (~ spl66_5 | ~ spl66_21), inference(avatar_split_clause, [], [f4824, f687, f619])).
fof(f619, plain, (spl66_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_5])])).
fof(f4824, plain, (~ (e0 = op(e3, e2)) | ~ spl66_21), inference(backward_demodulation, [], [f252, f689])).
fof(f252, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4823, plain, (~ spl66_19 | ~ spl66_27), inference(avatar_split_clause, [], [f4819, f712, f678])).
fof(f678, plain, (spl66_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_19])])).
fof(f712, plain, (spl66_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_27])])).
fof(f4819, plain, (~ (e2 = op(e2, e3)) | ~ spl66_27), inference(backward_demodulation, [], [f275, f714])).
fof(f714, plain, ((e2 = op(e2, e1)) | ~ spl66_27), inference(avatar_component_clause, [], [f712])).
fof(f275, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4822, plain, (~ spl66_11 | ~ spl66_27), inference(avatar_split_clause, [], [f4817, f712, f644])).
fof(f644, plain, (spl66_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_11])])).
fof(f4817, plain, (~ (e2 = op(e3, e1)) | ~ spl66_27), inference(backward_demodulation, [], [f246, f714])).
fof(f246, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4816, plain, (~ spl66_30 | ~ spl66_44 | ~ spl66_63 | spl66_168), inference(avatar_contradiction_clause, [], [f4815])).
fof(f4815, plain, ($false | (~ spl66_30 | ~ spl66_44 | ~ spl66_63 | spl66_168)), inference(subsumption_resolution, [], [f4811, f786])).
fof(f786, plain, ((e3 = op(e1, e1)) | ~ spl66_44), inference(avatar_component_clause, [], [f784])).
fof(f784, plain, (spl66_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_44])])).
fof(f4811, plain, (~ (e3 = op(e1, e1)) | (~ spl66_30 | ~ spl66_63 | spl66_168)), inference(backward_demodulation, [], [f4655, f727])).
fof(f727, plain, ((e1 = op(e2, e0)) | ~ spl66_30), inference(avatar_component_clause, [], [f725])).
fof(f725, plain, (spl66_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_30])])).
fof(f4655, plain, (~ (e3 = op(op(e2, e0), op(e2, e0))) | (~ spl66_63 | spl66_168)), inference(backward_demodulation, [], [f1622, f867])).
fof(f867, plain, ((op(e0, e0) = e2) | ~ spl66_63), inference(avatar_component_clause, [], [f865])).
fof(f865, plain, (spl66_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl66_63])])).
fof(f1622, plain, (~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0))) | spl66_168), inference(avatar_component_clause, [], [f1620])).
fof(f1620, plain, (spl66_168 <=> (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_168])])).
fof(f4814, plain, (spl66_25 | ~ spl66_30 | ~ spl66_69), inference(avatar_contradiction_clause, [], [f4813])).
fof(f4813, plain, ($false | (spl66_25 | ~ spl66_30 | ~ spl66_69)), inference(subsumption_resolution, [], [f4809, f705])).
fof(f705, plain, (~ (e0 = op(e2, e1)) | spl66_25), inference(avatar_component_clause, [], [f704])).
fof(f704, plain, (spl66_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_25])])).
fof(f4809, plain, ((e0 = op(e2, e1)) | (~ spl66_30 | ~ spl66_69)), inference(backward_demodulation, [], [f930, f727])).
fof(f930, plain, ((e0 = op(e2, op(e2, e0))) | ~ spl66_69), inference(avatar_component_clause, [], [f928])).
fof(f928, plain, (spl66_69 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_69])])).
fof(f4802, plain, (~ spl66_19 | ~ spl66_35), inference(avatar_split_clause, [], [f4798, f746, f678])).
fof(f746, plain, (spl66_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_35])])).
fof(f4798, plain, (~ (e2 = op(e2, e3)) | ~ spl66_35), inference(backward_demodulation, [], [f256, f748])).
fof(f748, plain, ((e2 = op(e1, e3)) | ~ spl66_35), inference(avatar_component_clause, [], [f746])).
fof(f256, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4791, plain, (~ spl66_33 | ~ spl66_45), inference(avatar_split_clause, [], [f4785, f789, f738])).
fof(f738, plain, (spl66_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_33])])).
fof(f789, plain, (spl66_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_45])])).
fof(f4785, plain, (~ (e0 = op(e1, e3)) | ~ spl66_45), inference(backward_demodulation, [], [f267, f791])).
fof(f791, plain, ((e0 = op(e1, e0)) | ~ spl66_45), inference(avatar_component_clause, [], [f789])).
fof(f267, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4781, plain, (~ spl66_2 | ~ spl66_50), inference(avatar_split_clause, [], [f4777, f810, f606])).
fof(f810, plain, (spl66_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_50])])).
fof(f4777, plain, (~ (e1 = op(e3, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f255, f812])).
fof(f812, plain, ((e1 = op(e0, e3)) | ~ spl66_50), inference(avatar_component_clause, [], [f810])).
fof(f255, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4774, plain, (~ spl66_52 | ~ spl66_56), inference(avatar_split_clause, [], [f4771, f835, f818])).
fof(f818, plain, (spl66_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_52])])).
fof(f835, plain, (spl66_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_56])])).
fof(f4771, plain, (~ (e3 = op(e0, e3)) | ~ spl66_56), inference(backward_demodulation, [], [f264, f837])).
fof(f837, plain, ((e3 = op(e0, e2)) | ~ spl66_56), inference(avatar_component_clause, [], [f835])).
fof(f264, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4773, plain, (~ spl66_24 | ~ spl66_56), inference(avatar_split_clause, [], [f4769, f835, f699])).
fof(f699, plain, (spl66_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_24])])).
fof(f4769, plain, (~ (e3 = op(e2, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f248, f837])).
fof(f248, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4767, plain, (spl66_1 | ~ spl66_16 | ~ spl66_146), inference(avatar_split_clause, [], [f4766, f1509, f665, f602])).
fof(f602, plain, (spl66_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_1])])).
fof(f665, plain, (spl66_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_16])])).
fof(f1509, plain, (spl66_146 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_146])])).
fof(f4766, plain, ((e0 = op(e3, e3)) | (~ spl66_16 | ~ spl66_146)), inference(forward_demodulation, [], [f1511, f667])).
fof(f667, plain, ((e3 = op(e3, e0)) | ~ spl66_16), inference(avatar_component_clause, [], [f665])).
fof(f1511, plain, ((e0 = op(e3, op(e3, e0))) | ~ spl66_146), inference(avatar_component_clause, [], [f1509])).
fof(f4765, plain, (~ spl66_11 | ~ spl66_44 | spl66_149), inference(avatar_split_clause, [], [f4764, f1527, f784, f644])).
fof(f1527, plain, (spl66_149 <=> (e2 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_149])])).
fof(f4764, plain, (~ (e2 = op(e3, e1)) | (~ spl66_44 | spl66_149)), inference(forward_demodulation, [], [f1529, f786])).
fof(f1529, plain, (~ (e2 = op(op(e1, e1), e1)) | spl66_149), inference(avatar_component_clause, [], [f1527])).
fof(f4752, plain, (~ spl66_48 | ~ spl66_16), inference(avatar_split_clause, [], [f4751, f665, f801])).
fof(f801, plain, (spl66_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_48])])).
fof(f4751, plain, (~ (e3 = op(e1, e0)) | ~ spl66_16), inference(forward_demodulation, [], [f239, f667])).
fof(f239, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f4750, plain, (~ spl66_47 | ~ spl66_63), inference(avatar_split_clause, [], [f4749, f865, f797])).
fof(f797, plain, (spl66_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_47])])).
fof(f4749, plain, (~ (e2 = op(e1, e0)) | ~ spl66_63), inference(forward_demodulation, [], [f235, f867])).
fof(f235, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f4748, plain, (~ spl66_36 | ~ spl66_44), inference(avatar_split_clause, [], [f4747, f784, f750])).
fof(f750, plain, (spl66_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_36])])).
fof(f4747, plain, (~ (e3 = op(e1, e3)) | ~ spl66_44), inference(forward_demodulation, [], [f269, f786])).
fof(f269, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4744, plain, (~ spl66_4 | ~ spl66_16), inference(avatar_split_clause, [], [f4743, f665, f614])).
fof(f614, plain, (spl66_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_4])])).
fof(f4743, plain, (~ (e3 = op(e3, e3)) | ~ spl66_16), inference(forward_demodulation, [], [f279, f667])).
fof(f279, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4738, plain, (spl66_6 | ~ spl66_11 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f4737])).
fof(f4737, plain, ($false | (spl66_6 | ~ spl66_11 | ~ spl66_145)), inference(subsumption_resolution, [], [f4735, f624])).
fof(f624, plain, (~ (e1 = op(e3, e2)) | spl66_6), inference(avatar_component_clause, [], [f623])).
fof(f623, plain, (spl66_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_6])])).
fof(f4735, plain, ((e1 = op(e3, e2)) | (~ spl66_11 | ~ spl66_145)), inference(backward_demodulation, [], [f1506, f646])).
fof(f646, plain, ((e2 = op(e3, e1)) | ~ spl66_11), inference(avatar_component_clause, [], [f644])).
fof(f1506, plain, ((e1 = op(e3, op(e3, e1))) | ~ spl66_145), inference(avatar_component_clause, [], [f1504])).
fof(f1504, plain, (spl66_145 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_145])])).
fof(f4736, plain, (~ spl66_7 | ~ spl66_11), inference(avatar_split_clause, [], [f4733, f644, f627])).
fof(f627, plain, (spl66_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_7])])).
fof(f4733, plain, (~ (e2 = op(e3, e2)) | ~ spl66_11), inference(backward_demodulation, [], [f280, f646])).
fof(f280, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4732, plain, (~ spl66_12 | ~ spl66_16), inference(avatar_split_clause, [], [f4729, f665, f648])).
fof(f648, plain, (spl66_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_12])])).
fof(f4729, plain, (~ (e3 = op(e3, e1)) | ~ spl66_16), inference(backward_demodulation, [], [f277, f667])).
fof(f277, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4725, plain, (spl66_19 | ~ spl66_24 | ~ spl66_67), inference(avatar_split_clause, [], [f4722, f918, f699, f678])).
fof(f918, plain, (spl66_67 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_67])])).
fof(f4722, plain, ((e2 = op(e2, e3)) | (~ spl66_24 | ~ spl66_67)), inference(backward_demodulation, [], [f920, f701])).
fof(f701, plain, ((e3 = op(e2, e2)) | ~ spl66_24), inference(avatar_component_clause, [], [f699])).
fof(f920, plain, ((e2 = op(e2, op(e2, e2))) | ~ spl66_67), inference(avatar_component_clause, [], [f918])).
fof(f4724, plain, (~ spl66_20 | ~ spl66_24), inference(avatar_split_clause, [], [f4721, f699, f682])).
fof(f682, plain, (spl66_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_20])])).
fof(f4721, plain, (~ (e3 = op(e2, e3)) | ~ spl66_24), inference(backward_demodulation, [], [f276, f701])).
fof(f276, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4719, plain, (~ spl66_22 | ~ spl66_26), inference(avatar_split_clause, [], [f4714, f708, f691])).
fof(f691, plain, (spl66_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_22])])).
fof(f708, plain, (spl66_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_26])])).
fof(f4714, plain, (~ (e1 = op(e2, e2)) | ~ spl66_26), inference(backward_demodulation, [], [f274, f710])).
fof(f710, plain, ((e1 = op(e2, e1)) | ~ spl66_26), inference(avatar_component_clause, [], [f708])).
fof(f274, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4709, plain, (~ spl66_25 | ~ spl66_29), inference(avatar_split_clause, [], [f4699, f721, f704])).
fof(f721, plain, (spl66_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_29])])).
fof(f4699, plain, (~ (e0 = op(e2, e1)) | ~ spl66_29), inference(backward_demodulation, [], [f271, f723])).
fof(f723, plain, ((e0 = op(e2, e0)) | ~ spl66_29), inference(avatar_component_clause, [], [f721])).
fof(f271, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4697, plain, (~ spl66_7 | ~ spl66_39), inference(avatar_split_clause, [], [f4694, f763, f627])).
fof(f763, plain, (spl66_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_39])])).
fof(f4694, plain, (~ (e2 = op(e3, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f251, f765])).
fof(f765, plain, ((e2 = op(e1, e2)) | ~ spl66_39), inference(avatar_component_clause, [], [f763])).
fof(f251, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4690, plain, (~ spl66_40 | ~ spl66_44), inference(avatar_split_clause, [], [f4685, f784, f767])).
fof(f767, plain, (spl66_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_40])])).
fof(f4685, plain, (~ (e3 = op(e1, e2)) | ~ spl66_44), inference(backward_demodulation, [], [f268, f786])).
fof(f268, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f4689, plain, (~ spl66_12 | ~ spl66_44), inference(avatar_split_clause, [], [f4684, f784, f648])).
fof(f4684, plain, (~ (e3 = op(e3, e1)) | ~ spl66_44), inference(backward_demodulation, [], [f245, f786])).
fof(f245, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4682, plain, (~ spl66_20 | ~ spl66_52), inference(avatar_split_clause, [], [f4680, f818, f682])).
fof(f4680, plain, (~ (e3 = op(e2, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f254, f820])).
fof(f820, plain, ((e3 = op(e0, e3)) | ~ spl66_52), inference(avatar_component_clause, [], [f818])).
fof(f254, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4678, plain, (~ spl66_22 | ~ spl66_54), inference(avatar_split_clause, [], [f4672, f827, f691])).
fof(f827, plain, (spl66_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_54])])).
fof(f4672, plain, (~ (e1 = op(e2, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f248, f829])).
fof(f829, plain, ((e1 = op(e0, e2)) | ~ spl66_54), inference(avatar_component_clause, [], [f827])).
fof(f4670, plain, (~ spl66_53 | ~ spl66_57), inference(avatar_split_clause, [], [f4663, f840, f823])).
fof(f840, plain, (spl66_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_57])])).
fof(f4663, plain, (~ (e0 = op(e0, e2)) | ~ spl66_57), inference(backward_demodulation, [], [f262, f842])).
fof(f842, plain, ((e0 = op(e0, e1)) | ~ spl66_57), inference(avatar_component_clause, [], [f840])).
fof(f262, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f4669, plain, (~ spl66_25 | ~ spl66_57), inference(avatar_split_clause, [], [f4661, f840, f704])).
fof(f4661, plain, (~ (e0 = op(e2, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f242, f842])).
fof(f242, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4657, plain, (~ spl66_51 | ~ spl66_63), inference(avatar_split_clause, [], [f4648, f865, f814])).
fof(f814, plain, (spl66_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_51])])).
fof(f4648, plain, (~ (e2 = op(e0, e3)) | ~ spl66_63), inference(backward_demodulation, [], [f261, f867])).
fof(f261, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4656, plain, (~ spl66_31 | ~ spl66_63), inference(avatar_split_clause, [], [f4644, f865, f729])).
fof(f729, plain, (spl66_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_31])])).
fof(f4644, plain, (~ (e2 = op(e2, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f236, f867])).
fof(f236, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f4643, plain, (~ spl66_12 | ~ spl66_2 | spl66_143), inference(avatar_split_clause, [], [f4642, f1494, f606, f648])).
fof(f1494, plain, (spl66_143 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_143])])).
fof(f4642, plain, (~ (e3 = op(e3, e1)) | (~ spl66_2 | spl66_143)), inference(forward_demodulation, [], [f1495, f608])).
fof(f608, plain, ((e1 = op(e3, e3)) | ~ spl66_2), inference(avatar_component_clause, [], [f606])).
fof(f1495, plain, (~ (e3 = op(e3, op(e3, e3))) | spl66_143), inference(avatar_component_clause, [], [f1494])).
fof(f4636, plain, (~ spl66_64 | ~ spl66_160 | spl66_170), inference(avatar_split_clause, [], [f4634, f1632, f1579, f869])).
fof(f869, plain, (spl66_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl66_64])])).
fof(f1632, plain, (spl66_170 <=> (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_170])])).
fof(f4634, plain, (~ (op(e0, e0) = e3) | (~ spl66_160 | spl66_170)), inference(backward_demodulation, [], [f1634, f1580])).
fof(f1634, plain, (~ (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2))) | spl66_170), inference(avatar_component_clause, [], [f1632])).
fof(f4631, plain, (~ spl66_63 | ~ spl66_162 | spl66_167), inference(avatar_split_clause, [], [f4517, f1614, f1589, f865])).
fof(f1589, plain, (spl66_162 <=> (e0 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_162])])).
fof(f1614, plain, (spl66_167 <=> (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_167])])).
fof(f4517, plain, (~ (op(e0, e0) = e2) | (~ spl66_162 | spl66_167)), inference(backward_demodulation, [], [f1616, f1590])).
fof(f1590, plain, ((e0 = op(op(e3, e3), e3)) | ~ spl66_162), inference(avatar_component_clause, [], [f1589])).
fof(f1616, plain, (~ (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3))) | spl66_167), inference(avatar_component_clause, [], [f1614])).
fof(f4623, plain, (~ spl66_14 | ~ spl66_46), inference(avatar_split_clause, [], [f4622, f793, f657])).
fof(f657, plain, (spl66_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_14])])).
fof(f793, plain, (spl66_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_46])])).
fof(f4622, plain, (~ (e1 = op(e3, e0)) | ~ spl66_46), inference(forward_demodulation, [], [f239, f795])).
fof(f795, plain, ((e1 = op(e1, e0)) | ~ spl66_46), inference(avatar_component_clause, [], [f793])).
fof(f4621, plain, (~ spl66_14 | ~ spl66_2), inference(avatar_split_clause, [], [f4620, f606, f657])).
fof(f4620, plain, (~ (e1 = op(e3, e0)) | ~ spl66_2), inference(forward_demodulation, [], [f279, f608])).
fof(f4613, plain, (~ spl66_3 | ~ spl66_7), inference(avatar_split_clause, [], [f4610, f627, f610])).
fof(f610, plain, (spl66_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_3])])).
fof(f4610, plain, (~ (e2 = op(e3, e3)) | ~ spl66_7), inference(backward_demodulation, [], [f282, f629])).
fof(f629, plain, ((e2 = op(e3, e2)) | ~ spl66_7), inference(avatar_component_clause, [], [f627])).
fof(f282, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4608, plain, (~ spl66_6 | ~ spl66_22), inference(avatar_split_clause, [], [f4601, f691, f623])).
fof(f4601, plain, (~ (e1 = op(e3, e2)) | ~ spl66_22), inference(backward_demodulation, [], [f252, f693])).
fof(f693, plain, ((e1 = op(e2, e2)) | ~ spl66_22), inference(avatar_component_clause, [], [f691])).
fof(f4600, plain, (~ spl66_21 | ~ spl66_25), inference(avatar_split_clause, [], [f4599, f704, f687])).
fof(f4599, plain, (~ (e0 = op(e2, e2)) | ~ spl66_25), inference(backward_demodulation, [], [f274, f706])).
fof(f706, plain, ((e0 = op(e2, e1)) | ~ spl66_25), inference(avatar_component_clause, [], [f704])).
fof(f4583, plain, (~ spl66_39 | ~ spl66_43), inference(avatar_split_clause, [], [f4578, f780, f763])).
fof(f780, plain, (spl66_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_43])])).
fof(f4578, plain, (~ (e2 = op(e1, e2)) | ~ spl66_43), inference(backward_demodulation, [], [f268, f782])).
fof(f782, plain, ((e2 = op(e1, e1)) | ~ spl66_43), inference(avatar_component_clause, [], [f780])).
fof(f4575, plain, (~ spl66_34 | ~ spl66_46), inference(avatar_split_clause, [], [f4571, f793, f742])).
fof(f742, plain, (spl66_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_34])])).
fof(f4571, plain, (~ (e1 = op(e1, e3)) | ~ spl66_46), inference(backward_demodulation, [], [f267, f795])).
fof(f4567, plain, (~ spl66_3 | ~ spl66_51), inference(avatar_split_clause, [], [f4564, f814, f610])).
fof(f4564, plain, (~ (e2 = op(e3, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f255, f816])).
fof(f816, plain, ((e2 = op(e0, e3)) | ~ spl66_51), inference(avatar_component_clause, [], [f814])).
fof(f4562, plain, (~ spl66_49 | ~ spl66_53), inference(avatar_split_clause, [], [f4556, f823, f806])).
fof(f806, plain, (spl66_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_49])])).
fof(f4556, plain, (~ (e0 = op(e0, e3)) | ~ spl66_53), inference(backward_demodulation, [], [f264, f825])).
fof(f825, plain, ((e0 = op(e0, e2)) | ~ spl66_53), inference(avatar_component_clause, [], [f823])).
fof(f4561, plain, (~ spl66_21 | ~ spl66_53), inference(avatar_split_clause, [], [f4554, f823, f687])).
fof(f4554, plain, (~ (e0 = op(e2, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f248, f825])).
fof(f4552, plain, (~ spl66_26 | ~ spl66_58), inference(avatar_split_clause, [], [f4546, f844, f708])).
fof(f844, plain, (spl66_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_58])])).
fof(f4546, plain, (~ (e1 = op(e2, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f242, f846])).
fof(f846, plain, ((e1 = op(e0, e1)) | ~ spl66_58), inference(avatar_component_clause, [], [f844])).
fof(f4538, plain, (~ spl66_56 | ~ spl66_64), inference(avatar_split_clause, [], [f4531, f869, f835])).
fof(f4531, plain, (~ (e3 = op(e0, e2)) | ~ spl66_64), inference(backward_demodulation, [], [f260, f871])).
fof(f871, plain, ((op(e0, e0) = e3) | ~ spl66_64), inference(avatar_component_clause, [], [f869])).
fof(f260, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f4537, plain, (~ spl66_48 | ~ spl66_64), inference(avatar_split_clause, [], [f4529, f869, f801])).
fof(f4529, plain, (~ (e3 = op(e1, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f235, f871])).
fof(f4526, plain, (spl66_2 | ~ spl66_12 | ~ spl66_145), inference(avatar_split_clause, [], [f4525, f1504, f648, f606])).
fof(f4525, plain, ((e1 = op(e3, e3)) | (~ spl66_12 | ~ spl66_145)), inference(forward_demodulation, [], [f1506, f650])).
fof(f650, plain, ((e3 = op(e3, e1)) | ~ spl66_12), inference(avatar_component_clause, [], [f648])).
fof(f4516, plain, (~ spl66_64 | ~ spl66_166 | spl66_169), inference(avatar_split_clause, [], [f4514, f1626, f1609, f869])).
fof(f1609, plain, (spl66_166 <=> (e0 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_166])])).
fof(f1626, plain, (spl66_169 <=> (e3 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_169])])).
fof(f4514, plain, (~ (op(e0, e0) = e3) | (~ spl66_166 | spl66_169)), inference(backward_demodulation, [], [f1628, f1610])).
fof(f1610, plain, ((e0 = op(op(e1, e1), e1)) | ~ spl66_166), inference(avatar_component_clause, [], [f1609])).
fof(f1628, plain, (~ (e3 = op(op(op(e1, e1), e1), op(op(e1, e1), e1))) | spl66_169), inference(avatar_component_clause, [], [f1626])).
fof(f4511, plain, (~ spl66_2 | spl66_156 | ~ spl66_157), inference(avatar_split_clause, [], [f4510, f1564, f1560, f606])).
fof(f1560, plain, (spl66_156 <=> (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_156])])).
fof(f1564, plain, (spl66_157 <=> (e3 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_157])])).
fof(f4510, plain, (~ (e1 = op(e3, e3)) | (spl66_156 | ~ spl66_157)), inference(backward_demodulation, [], [f1562, f1565])).
fof(f1565, plain, ((e3 = op(op(e0, e0), e0)) | ~ spl66_157), inference(avatar_component_clause, [], [f1564])).
fof(f1562, plain, (~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0))) | spl66_156), inference(avatar_component_clause, [], [f1560])).
fof(f4499, plain, (~ spl66_23 | ~ spl66_31), inference(avatar_split_clause, [], [f4498, f729, f695])).
fof(f695, plain, (spl66_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_23])])).
fof(f4498, plain, (~ (e2 = op(e2, e2)) | ~ spl66_31), inference(forward_demodulation, [], [f272, f731])).
fof(f731, plain, ((e2 = op(e2, e0)) | ~ spl66_31), inference(avatar_component_clause, [], [f729])).
fof(f272, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4497, plain, (~ spl66_2 | ~ spl66_151 | spl66_159), inference(avatar_split_clause, [], [f4496, f1574, f1536, f606])).
fof(f1536, plain, (spl66_151 <=> (e3 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_151])])).
fof(f1574, plain, (spl66_159 <=> (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_159])])).
fof(f4496, plain, (~ (e1 = op(e3, e3)) | (~ spl66_151 | spl66_159)), inference(forward_demodulation, [], [f1576, f1537])).
fof(f1537, plain, ((e3 = op(op(e2, e2), e2)) | ~ spl66_151), inference(avatar_component_clause, [], [f1536])).
fof(f1576, plain, (~ (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2))) | spl66_159), inference(avatar_component_clause, [], [f1574])).
fof(f4494, plain, (~ spl66_4 | ~ spl66_20), inference(avatar_split_clause, [], [f4493, f682, f614])).
fof(f4493, plain, (~ (e3 = op(e3, e3)) | ~ spl66_20), inference(forward_demodulation, [], [f258, f684])).
fof(f684, plain, ((e3 = op(e2, e3)) | ~ spl66_20), inference(avatar_component_clause, [], [f682])).
fof(f258, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4492, plain, (~ spl66_4 | ~ spl66_12), inference(avatar_split_clause, [], [f4491, f648, f614])).
fof(f4491, plain, (~ (e3 = op(e3, e3)) | ~ spl66_12), inference(forward_demodulation, [], [f281, f650])).
fof(f4478, plain, (~ spl66_3 | ~ spl66_48 | ~ spl66_62 | spl66_163), inference(avatar_contradiction_clause, [], [f4477])).
fof(f4477, plain, ($false | (~ spl66_3 | ~ spl66_48 | ~ spl66_62 | spl66_163)), inference(subsumption_resolution, [], [f4476, f612])).
fof(f612, plain, ((e2 = op(e3, e3)) | ~ spl66_3), inference(avatar_component_clause, [], [f610])).
fof(f4476, plain, (~ (e2 = op(e3, e3)) | (~ spl66_48 | ~ spl66_62 | spl66_163)), inference(forward_demodulation, [], [f4475, f803])).
fof(f803, plain, ((e3 = op(e1, e0)) | ~ spl66_48), inference(avatar_component_clause, [], [f801])).
fof(f4475, plain, (~ (e2 = op(op(e1, e0), op(e1, e0))) | (~ spl66_62 | spl66_163)), inference(forward_demodulation, [], [f1596, f863])).
fof(f863, plain, ((op(e0, e0) = e1) | ~ spl66_62), inference(avatar_component_clause, [], [f861])).
fof(f861, plain, (spl66_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl66_62])])).
fof(f1596, plain, (~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0))) | spl66_163), inference(avatar_component_clause, [], [f1594])).
fof(f1594, plain, (spl66_163 <=> (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_163])])).
fof(f4471, plain, (~ spl66_8 | ~ spl66_12), inference(avatar_split_clause, [], [f4469, f648, f631])).
fof(f631, plain, (spl66_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_8])])).
fof(f4469, plain, (~ (e3 = op(e3, e2)) | ~ spl66_12), inference(backward_demodulation, [], [f280, f650])).
fof(f4464, plain, (~ spl66_18 | ~ spl66_26), inference(avatar_split_clause, [], [f4459, f708, f674])).
fof(f674, plain, (spl66_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_18])])).
fof(f4459, plain, (~ (e1 = op(e2, e3)) | ~ spl66_26), inference(backward_demodulation, [], [f275, f710])).
fof(f4463, plain, (~ spl66_10 | ~ spl66_26), inference(avatar_split_clause, [], [f4458, f708, f640])).
fof(f4458, plain, (~ (e1 = op(e3, e1)) | ~ spl66_26), inference(backward_demodulation, [], [f246, f710])).
fof(f4457, plain, (~ spl66_27 | ~ spl66_31), inference(avatar_split_clause, [], [f4453, f729, f712])).
fof(f4453, plain, (~ (e2 = op(e2, e1)) | ~ spl66_31), inference(backward_demodulation, [], [f271, f731])).
fof(f4451, plain, (~ spl66_18 | ~ spl66_34), inference(avatar_split_clause, [], [f4447, f742, f674])).
fof(f4447, plain, (~ (e1 = op(e2, e3)) | ~ spl66_34), inference(backward_demodulation, [], [f256, f744])).
fof(f744, plain, ((e1 = op(e1, e3)) | ~ spl66_34), inference(avatar_component_clause, [], [f742])).
fof(f4446, plain, (~ spl66_33 | ~ spl66_41), inference(avatar_split_clause, [], [f4440, f772, f738])).
fof(f772, plain, (spl66_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_41])])).
fof(f4440, plain, (~ (e0 = op(e1, e3)) | ~ spl66_41), inference(backward_demodulation, [], [f269, f774])).
fof(f774, plain, ((e0 = op(e1, e1)) | ~ spl66_41), inference(avatar_component_clause, [], [f772])).
fof(f4435, plain, (~ spl66_44 | ~ spl66_48), inference(avatar_split_clause, [], [f4431, f801, f784])).
fof(f4431, plain, (~ (e3 = op(e1, e1)) | ~ spl66_48), inference(backward_demodulation, [], [f265, f803])).
fof(f265, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f4434, plain, (~ spl66_32 | ~ spl66_48), inference(avatar_split_clause, [], [f4430, f801, f733])).
fof(f733, plain, (spl66_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_32])])).
fof(f4430, plain, (~ (e3 = op(e2, e0)) | ~ spl66_48), inference(backward_demodulation, [], [f238, f803])).
fof(f238, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f4428, plain, (~ spl66_33 | ~ spl66_49), inference(avatar_split_clause, [], [f4422, f806, f738])).
fof(f4422, plain, (~ (e0 = op(e1, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f253, f808])).
fof(f808, plain, ((e0 = op(e0, e3)) | ~ spl66_49), inference(avatar_component_clause, [], [f806])).
fof(f253, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4418, plain, (~ spl66_8 | ~ spl66_56), inference(avatar_split_clause, [], [f4413, f835, f631])).
fof(f4413, plain, (~ (e3 = op(e3, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f249, f837])).
fof(f249, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4411, plain, (~ spl66_27 | ~ spl66_59), inference(avatar_split_clause, [], [f4405, f848, f712])).
fof(f848, plain, (spl66_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_59])])).
fof(f4405, plain, (~ (e2 = op(e2, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f242, f850])).
fof(f850, plain, ((e2 = op(e0, e1)) | ~ spl66_59), inference(avatar_component_clause, [], [f848])).
fof(f4402, plain, (~ spl66_48 | ~ spl66_62 | spl66_157), inference(avatar_split_clause, [], [f4396, f1564, f861, f801])).
fof(f4396, plain, (~ (e3 = op(e1, e0)) | (~ spl66_62 | spl66_157)), inference(backward_demodulation, [], [f1566, f863])).
fof(f1566, plain, (~ (e3 = op(op(e0, e0), e0)) | spl66_157), inference(avatar_component_clause, [], [f1564])).
fof(f4400, plain, (~ spl66_54 | ~ spl66_62), inference(avatar_split_clause, [], [f4390, f861, f827])).
fof(f4390, plain, (~ (e1 = op(e0, e2)) | ~ spl66_62), inference(backward_demodulation, [], [f260, f863])).
fof(f4399, plain, (~ spl66_46 | ~ spl66_62), inference(avatar_split_clause, [], [f4387, f861, f793])).
fof(f4387, plain, (~ (e1 = op(e1, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f235, f863])).
fof(f4378, plain, (~ spl66_56 | ~ spl66_21 | spl66_151), inference(avatar_split_clause, [], [f4377, f1536, f687, f835])).
fof(f4377, plain, (~ (e3 = op(e0, e2)) | (~ spl66_21 | spl66_151)), inference(forward_demodulation, [], [f1538, f689])).
fof(f1538, plain, (~ (e3 = op(op(e2, e2), e2)) | spl66_151), inference(avatar_component_clause, [], [f1536])).
fof(f4376, plain, (~ spl66_18 | ~ spl66_3 | spl66_155), inference(avatar_split_clause, [], [f4375, f1555, f610, f674])).
fof(f1555, plain, (spl66_155 <=> (e1 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_155])])).
fof(f4375, plain, (~ (e1 = op(e2, e3)) | (~ spl66_3 | spl66_155)), inference(forward_demodulation, [], [f1557, f612])).
fof(f1557, plain, (~ (e1 = op(op(e3, e3), e3)) | spl66_155), inference(avatar_component_clause, [], [f1555])).
fof(f4367, plain, (~ spl66_54 | ~ spl66_21 | spl66_152), inference(avatar_split_clause, [], [f4327, f1541, f687, f827])).
fof(f1541, plain, (spl66_152 <=> (e1 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_152])])).
fof(f4327, plain, (~ (e1 = op(e0, e2)) | (~ spl66_21 | spl66_152)), inference(backward_demodulation, [], [f1543, f689])).
fof(f1543, plain, (~ (e1 = op(op(e2, e2), e2)) | spl66_152), inference(avatar_component_clause, [], [f1541])).
fof(f4361, plain, (spl66_31 | ~ spl66_21 | ~ spl66_67), inference(avatar_split_clause, [], [f4325, f918, f687, f729])).
fof(f4325, plain, ((e2 = op(e2, e0)) | (~ spl66_21 | ~ spl66_67)), inference(backward_demodulation, [], [f920, f689])).
fof(f4359, plain, (~ spl66_11 | ~ spl66_3), inference(avatar_split_clause, [], [f4358, f610, f644])).
fof(f4358, plain, (~ (e2 = op(e3, e1)) | ~ spl66_3), inference(forward_demodulation, [], [f281, f612])).
fof(f4355, plain, (~ spl66_5 | ~ spl66_13), inference(avatar_split_clause, [], [f4350, f653, f619])).
fof(f653, plain, (spl66_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_13])])).
fof(f4350, plain, (~ (e0 = op(e3, e2)) | ~ spl66_13), inference(backward_demodulation, [], [f278, f655])).
fof(f655, plain, ((e0 = op(e3, e0)) | ~ spl66_13), inference(avatar_component_clause, [], [f653])).
fof(f278, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4340, plain, (~ spl66_21 | ~ spl66_44 | ~ spl66_54 | spl66_170), inference(avatar_contradiction_clause, [], [f4339])).
fof(f4339, plain, ($false | (~ spl66_21 | ~ spl66_44 | ~ spl66_54 | spl66_170)), inference(subsumption_resolution, [], [f4338, f786])).
fof(f4338, plain, (~ (e3 = op(e1, e1)) | (~ spl66_21 | ~ spl66_54 | spl66_170)), inference(forward_demodulation, [], [f4330, f829])).
fof(f4330, plain, (~ (e3 = op(op(e0, e2), op(e0, e2))) | (~ spl66_21 | spl66_170)), inference(backward_demodulation, [], [f1634, f689])).
fof(f4317, plain, (~ spl66_20 | ~ spl66_32), inference(avatar_split_clause, [], [f4313, f733, f682])).
fof(f4313, plain, (~ (e3 = op(e2, e3)) | ~ spl66_32), inference(backward_demodulation, [], [f273, f735])).
fof(f735, plain, ((e3 = op(e2, e0)) | ~ spl66_32), inference(avatar_component_clause, [], [f733])).
fof(f273, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4316, plain, (~ spl66_16 | ~ spl66_32), inference(avatar_split_clause, [], [f4311, f733, f665])).
fof(f4311, plain, (~ (e3 = op(e3, e0)) | ~ spl66_32), inference(backward_demodulation, [], [f240, f735])).
fof(f240, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f4260, plain, (~ spl66_49 | ~ spl66_57), inference(avatar_split_clause, [], [f4254, f840, f806])).
fof(f4254, plain, (~ (e0 = op(e0, e3)) | ~ spl66_57), inference(backward_demodulation, [], [f263, f842])).
fof(f263, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4259, plain, (~ spl66_41 | ~ spl66_57), inference(avatar_split_clause, [], [f4252, f840, f772])).
fof(f4252, plain, (~ (e0 = op(e1, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f241, f842])).
fof(f241, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f4251, plain, (~ spl66_32 | ~ spl66_63 | spl66_157), inference(avatar_split_clause, [], [f4246, f1564, f865, f733])).
fof(f4246, plain, (~ (e3 = op(e2, e0)) | (~ spl66_63 | spl66_157)), inference(backward_demodulation, [], [f1566, f867])).
fof(f4250, plain, (~ spl66_55 | ~ spl66_63), inference(avatar_split_clause, [], [f4241, f865, f831])).
fof(f831, plain, (spl66_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_55])])).
fof(f4241, plain, (~ (e2 = op(e0, e2)) | ~ spl66_63), inference(backward_demodulation, [], [f260, f867])).
fof(f4234, plain, (spl66_8 | ~ spl66_3 | ~ spl66_143), inference(avatar_split_clause, [], [f4233, f1494, f610, f631])).
fof(f4233, plain, ((e3 = op(e3, e2)) | (~ spl66_3 | ~ spl66_143)), inference(forward_demodulation, [], [f1496, f612])).
fof(f1496, plain, ((e3 = op(e3, op(e3, e3))) | ~ spl66_143), inference(avatar_component_clause, [], [f1494])).
fof(f4218, plain, (~ spl66_23 | ~ spl66_27), inference(avatar_split_clause, [], [f4217, f712, f695])).
fof(f4217, plain, (~ (e2 = op(e2, e2)) | ~ spl66_27), inference(forward_demodulation, [], [f274, f714])).
fof(f4205, plain, (~ spl66_1 | ~ spl66_5), inference(avatar_split_clause, [], [f4200, f619, f602])).
fof(f4200, plain, (~ (e0 = op(e3, e3)) | ~ spl66_5), inference(backward_demodulation, [], [f282, f621])).
fof(f621, plain, ((e0 = op(e3, e2)) | ~ spl66_5), inference(avatar_component_clause, [], [f619])).
fof(f4198, plain, (~ spl66_8 | ~ spl66_16), inference(avatar_split_clause, [], [f4195, f665, f631])).
fof(f4195, plain, (~ (e3 = op(e3, e2)) | ~ spl66_16), inference(backward_demodulation, [], [f278, f667])).
fof(f4188, plain, (~ spl66_1 | ~ spl66_22 | ~ spl66_40 | spl66_150), inference(avatar_split_clause, [], [f4185, f1532, f767, f691, f602])).
fof(f1532, plain, (spl66_150 <=> (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_150])])).
fof(f4185, plain, (~ (e0 = op(e3, e3)) | (~ spl66_22 | ~ spl66_40 | spl66_150)), inference(backward_demodulation, [], [f4140, f769])).
fof(f769, plain, ((e3 = op(e1, e2)) | ~ spl66_40), inference(avatar_component_clause, [], [f767])).
fof(f4140, plain, (~ (e0 = op(op(e1, e2), op(e1, e2))) | (~ spl66_22 | spl66_150)), inference(forward_demodulation, [], [f1534, f693])).
fof(f1534, plain, (~ (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2))) | spl66_150), inference(avatar_component_clause, [], [f1532])).
fof(f4186, plain, (~ spl66_8 | ~ spl66_40), inference(avatar_split_clause, [], [f4181, f767, f631])).
fof(f4181, plain, (~ (e3 = op(e3, e2)) | ~ spl66_40), inference(backward_demodulation, [], [f251, f769])).
fof(f4180, plain, (~ spl66_37 | ~ spl66_41), inference(avatar_split_clause, [], [f4172, f772, f755])).
fof(f755, plain, (spl66_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_37])])).
fof(f4172, plain, (~ (e0 = op(e1, e2)) | ~ spl66_41), inference(backward_demodulation, [], [f268, f774])).
fof(f4169, plain, (~ spl66_43 | ~ spl66_47), inference(avatar_split_clause, [], [f4163, f797, f780])).
fof(f4163, plain, (~ (e2 = op(e1, e1)) | ~ spl66_47), inference(backward_demodulation, [], [f265, f799])).
fof(f799, plain, ((e2 = op(e1, e0)) | ~ spl66_47), inference(avatar_component_clause, [], [f797])).
fof(f4160, plain, (~ spl66_1 | ~ spl66_49), inference(avatar_split_clause, [], [f4155, f806, f602])).
fof(f4155, plain, (~ (e0 = op(e3, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f255, f808])).
fof(f4153, plain, (~ spl66_52 | ~ spl66_60), inference(avatar_split_clause, [], [f4148, f852, f818])).
fof(f852, plain, (spl66_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_60])])).
fof(f4148, plain, (~ (e3 = op(e0, e3)) | ~ spl66_60), inference(backward_demodulation, [], [f263, f854])).
fof(f854, plain, ((e3 = op(e0, e1)) | ~ spl66_60), inference(avatar_component_clause, [], [f852])).
fof(f4152, plain, (spl66_50 | ~ spl66_60 | ~ spl66_78), inference(avatar_contradiction_clause, [], [f4151])).
fof(f4151, plain, ($false | (spl66_50 | ~ spl66_60 | ~ spl66_78)), inference(subsumption_resolution, [], [f4147, f811])).
fof(f811, plain, (~ (e1 = op(e0, e3)) | spl66_50), inference(avatar_component_clause, [], [f810])).
fof(f4147, plain, ((e1 = op(e0, e3)) | (~ spl66_60 | ~ spl66_78)), inference(backward_demodulation, [], [f979, f854])).
fof(f979, plain, ((e1 = op(e0, op(e0, e1))) | ~ spl66_78), inference(avatar_component_clause, [], [f977])).
fof(f977, plain, (spl66_78 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_78])])).
fof(f4150, plain, (~ spl66_28 | ~ spl66_60), inference(avatar_split_clause, [], [f4146, f852, f716])).
fof(f716, plain, (spl66_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_28])])).
fof(f4146, plain, (~ (e3 = op(e2, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f242, f854])).
fof(f4145, plain, (spl66_27 | ~ spl66_22 | ~ spl66_67), inference(avatar_split_clause, [], [f4144, f918, f691, f712])).
fof(f4144, plain, ((e2 = op(e2, e1)) | (~ spl66_22 | ~ spl66_67)), inference(forward_demodulation, [], [f920, f693])).
fof(f4139, plain, (spl66_40 | ~ spl66_22 | ~ spl66_151), inference(avatar_split_clause, [], [f4138, f1536, f691, f767])).
fof(f4138, plain, ((e3 = op(e1, e2)) | (~ spl66_22 | ~ spl66_151)), inference(forward_demodulation, [], [f1537, f693])).
fof(f4131, plain, (~ spl66_37 | ~ spl66_22 | spl66_160), inference(avatar_split_clause, [], [f4130, f1579, f691, f755])).
fof(f4130, plain, (~ (e0 = op(e1, e2)) | (~ spl66_22 | spl66_160)), inference(forward_demodulation, [], [f1581, f693])).
fof(f1581, plain, (~ (e0 = op(op(e2, e2), e2)) | spl66_160), inference(avatar_component_clause, [], [f1579])).
fof(f4116, plain, (~ spl66_3 | ~ spl66_148 | spl66_165), inference(avatar_split_clause, [], [f4115, f1604, f1522, f610])).
fof(f1522, plain, (spl66_148 <=> (e3 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_148])])).
fof(f1604, plain, (spl66_165 <=> (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_165])])).
fof(f4115, plain, (~ (e2 = op(e3, e3)) | (~ spl66_148 | spl66_165)), inference(forward_demodulation, [], [f1606, f1523])).
fof(f1523, plain, ((e3 = op(op(e1, e1), e1)) | ~ spl66_148), inference(avatar_component_clause, [], [f1522])).
fof(f1606, plain, (~ (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1))) | spl66_165), inference(avatar_component_clause, [], [f1604])).
fof(f4109, plain, (~ spl66_1 | ~ spl66_8 | ~ spl66_144), inference(avatar_contradiction_clause, [], [f4108])).
fof(f4108, plain, ($false | (~ spl66_1 | ~ spl66_8 | ~ spl66_144)), inference(subsumption_resolution, [], [f4107, f284])).
fof(f284, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax4)).
fof(f4107, plain, ((e0 = e2) | (~ spl66_1 | ~ spl66_8 | ~ spl66_144)), inference(forward_demodulation, [], [f4106, f604])).
fof(f604, plain, ((e0 = op(e3, e3)) | ~ spl66_1), inference(avatar_component_clause, [], [f602])).
fof(f4106, plain, ((e2 = op(e3, e3)) | (~ spl66_8 | ~ spl66_144)), inference(backward_demodulation, [], [f1501, f633])).
fof(f633, plain, ((e3 = op(e3, e2)) | ~ spl66_8), inference(avatar_component_clause, [], [f631])).
fof(f1501, plain, ((e2 = op(e3, op(e3, e2))) | ~ spl66_144), inference(avatar_component_clause, [], [f1499])).
fof(f1499, plain, (spl66_144 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_144])])).
fof(f4101, plain, (~ spl66_7 | ~ spl66_15), inference(avatar_split_clause, [], [f4098, f661, f627])).
fof(f661, plain, (spl66_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_15])])).
fof(f4098, plain, (~ (e2 = op(e3, e2)) | ~ spl66_15), inference(backward_demodulation, [], [f278, f663])).
fof(f663, plain, ((e2 = op(e3, e0)) | ~ spl66_15), inference(avatar_component_clause, [], [f661])).
fof(f4096, plain, (~ spl66_1 | ~ spl66_28 | ~ spl66_43 | spl66_147), inference(avatar_contradiction_clause, [], [f4095])).
fof(f4095, plain, ($false | (~ spl66_1 | ~ spl66_28 | ~ spl66_43 | spl66_147)), inference(subsumption_resolution, [], [f4091, f604])).
fof(f4091, plain, (~ (e0 = op(e3, e3)) | (~ spl66_28 | ~ spl66_43 | spl66_147)), inference(backward_demodulation, [], [f4030, f718])).
fof(f718, plain, ((e3 = op(e2, e1)) | ~ spl66_28), inference(avatar_component_clause, [], [f716])).
fof(f4030, plain, (~ (e0 = op(op(e2, e1), op(e2, e1))) | (~ spl66_43 | spl66_147)), inference(forward_demodulation, [], [f1520, f782])).
fof(f1520, plain, (~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1))) | spl66_147), inference(avatar_component_clause, [], [f1518])).
fof(f1518, plain, (spl66_147 <=> (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_147])])).
fof(f4092, plain, (~ spl66_12 | ~ spl66_28), inference(avatar_split_clause, [], [f4087, f716, f648])).
fof(f4087, plain, (~ (e3 = op(e3, e1)) | ~ spl66_28), inference(backward_demodulation, [], [f246, f718])).
fof(f4060, plain, (~ spl66_7 | ~ spl66_55), inference(avatar_split_clause, [], [f4057, f831, f627])).
fof(f4057, plain, (~ (e2 = op(e3, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f249, f833])).
fof(f833, plain, ((e2 = op(e0, e2)) | ~ spl66_55), inference(avatar_component_clause, [], [f831])).
fof(f4046, plain, (spl66_57 | ~ spl66_62 | ~ spl66_79), inference(avatar_split_clause, [], [f4038, f982, f861, f840])).
fof(f982, plain, (spl66_79 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_79])])).
fof(f4038, plain, ((e0 = op(e0, e1)) | (~ spl66_62 | ~ spl66_79)), inference(backward_demodulation, [], [f984, f863])).
fof(f984, plain, ((e0 = op(e0, op(e0, e0))) | ~ spl66_79), inference(avatar_component_clause, [], [f982])).
fof(f4045, plain, (~ spl66_58 | ~ spl66_62), inference(avatar_split_clause, [], [f4034, f861, f844])).
fof(f4034, plain, (~ (e1 = op(e0, e1)) | ~ spl66_62), inference(backward_demodulation, [], [f259, f863])).
fof(f259, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f4023, plain, (~ spl66_25 | ~ spl66_43 | spl66_166), inference(avatar_split_clause, [], [f4022, f1609, f780, f704])).
fof(f4022, plain, (~ (e0 = op(e2, e1)) | (~ spl66_43 | spl66_166)), inference(forward_demodulation, [], [f1611, f782])).
fof(f1611, plain, (~ (e0 = op(op(e1, e1), e1)) | spl66_166), inference(avatar_component_clause, [], [f1609])).
fof(f4017, plain, (~ spl66_59 | ~ spl66_43), inference(avatar_split_clause, [], [f4016, f780, f848])).
fof(f4016, plain, (~ (e2 = op(e0, e1)) | ~ spl66_43), inference(forward_demodulation, [], [f241, f782])).
fof(f4010, plain, (~ spl66_40 | ~ spl66_22 | spl66_151), inference(avatar_split_clause, [], [f3971, f1536, f691, f767])).
fof(f3971, plain, (~ (e3 = op(e1, e2)) | (~ spl66_22 | spl66_151)), inference(backward_demodulation, [], [f1538, f693])).
fof(f3985, plain, (~ spl66_3 | ~ spl66_19), inference(avatar_split_clause, [], [f3982, f678, f610])).
fof(f3982, plain, (~ (e2 = op(e3, e3)) | ~ spl66_19), inference(backward_demodulation, [], [f258, f680])).
fof(f680, plain, ((e2 = op(e2, e3)) | ~ spl66_19), inference(avatar_component_clause, [], [f678])).
fof(f3975, plain, (~ spl66_18 | ~ spl66_22), inference(avatar_split_clause, [], [f3966, f691, f674])).
fof(f3966, plain, (~ (e1 = op(e2, e3)) | ~ spl66_22), inference(backward_demodulation, [], [f276, f693])).
fof(f3955, plain, (~ spl66_24 | ~ spl66_32), inference(avatar_split_clause, [], [f3949, f733, f699])).
fof(f3949, plain, (~ (e3 = op(e2, e2)) | ~ spl66_32), inference(backward_demodulation, [], [f272, f735])).
fof(f3940, plain, (~ spl66_24 | ~ spl66_40), inference(avatar_split_clause, [], [f3936, f767, f699])).
fof(f3936, plain, (~ (e3 = op(e2, e2)) | ~ spl66_40), inference(backward_demodulation, [], [f250, f769])).
fof(f250, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3933, plain, (~ spl66_27 | ~ spl66_43), inference(avatar_split_clause, [], [f3927, f780, f712])).
fof(f3927, plain, (~ (e2 = op(e2, e1)) | ~ spl66_43), inference(backward_demodulation, [], [f244, f782])).
fof(f244, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3923, plain, (~ spl66_41 | ~ spl66_45), inference(avatar_split_clause, [], [f3915, f789, f772])).
fof(f3915, plain, (~ (e0 = op(e1, e1)) | ~ spl66_45), inference(backward_demodulation, [], [f265, f791])).
fof(f3913, plain, (~ spl66_5 | ~ spl66_53), inference(avatar_split_clause, [], [f3908, f823, f619])).
fof(f3908, plain, (~ (e0 = op(e3, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f249, f825])).
fof(f3905, plain, (spl66_53 | ~ spl66_63 | ~ spl66_79), inference(avatar_split_clause, [], [f3897, f982, f865, f823])).
fof(f3897, plain, ((e0 = op(e0, e2)) | (~ spl66_63 | ~ spl66_79)), inference(backward_demodulation, [], [f984, f867])).
fof(f3884, plain, (~ spl66_30 | ~ spl66_14), inference(avatar_split_clause, [], [f3883, f657, f725])).
fof(f3883, plain, (~ (e1 = op(e2, e0)) | ~ spl66_14), inference(forward_demodulation, [], [f240, f659])).
fof(f659, plain, ((e1 = op(e3, e0)) | ~ spl66_14), inference(avatar_component_clause, [], [f657])).
fof(f3851, plain, (~ spl66_2 | ~ spl66_18), inference(avatar_split_clause, [], [f3847, f674, f606])).
fof(f3847, plain, (~ (e1 = op(e3, e3)) | ~ spl66_18), inference(backward_demodulation, [], [f258, f676])).
fof(f676, plain, ((e1 = op(e2, e3)) | ~ spl66_18), inference(avatar_component_clause, [], [f674])).
fof(f3845, plain, (~ spl66_8 | ~ spl66_24), inference(avatar_split_clause, [], [f3841, f699, f631])).
fof(f3841, plain, (~ (e3 = op(e3, e2)) | ~ spl66_24), inference(backward_demodulation, [], [f252, f701])).
fof(f3840, plain, (~ spl66_17 | ~ spl66_29), inference(avatar_split_clause, [], [f3837, f721, f670])).
fof(f670, plain, (spl66_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_17])])).
fof(f3837, plain, (~ (e0 = op(e2, e3)) | ~ spl66_29), inference(backward_demodulation, [], [f273, f723])).
fof(f3833, plain, (~ spl66_17 | ~ spl66_33), inference(avatar_split_clause, [], [f3828, f738, f670])).
fof(f3828, plain, (~ (e0 = op(e2, e3)) | ~ spl66_33), inference(backward_demodulation, [], [f256, f740])).
fof(f740, plain, ((e0 = op(e1, e3)) | ~ spl66_33), inference(avatar_component_clause, [], [f738])).
fof(f3824, plain, (~ spl66_9 | ~ spl66_41), inference(avatar_split_clause, [], [f3816, f772, f636])).
fof(f636, plain, (spl66_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_9])])).
fof(f3816, plain, (~ (e0 = op(e3, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f245, f774])).
fof(f3804, plain, (~ spl66_23 | ~ spl66_55), inference(avatar_split_clause, [], [f3800, f831, f695])).
fof(f3800, plain, (~ (e2 = op(e2, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f248, f833])).
fof(f3795, plain, (~ spl66_53 | ~ spl66_61), inference(avatar_split_clause, [], [f3786, f857, f823])).
fof(f857, plain, (spl66_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_61])])).
fof(f3786, plain, (~ (e0 = op(e0, e2)) | ~ spl66_61), inference(backward_demodulation, [], [f260, f859])).
fof(f859, plain, ((e0 = op(e0, e0)) | ~ spl66_61), inference(avatar_component_clause, [], [f857])).
fof(f3793, plain, (~ spl66_45 | ~ spl66_61), inference(avatar_split_clause, [], [f3783, f857, f789])).
fof(f3783, plain, (~ (e0 = op(e1, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f235, f859])).
fof(f3768, plain, (~ spl66_4 | ~ spl66_52), inference(avatar_split_clause, [], [f3767, f818, f614])).
fof(f3767, plain, (~ (e3 = op(e3, e3)) | ~ spl66_52), inference(forward_demodulation, [], [f255, f820])).
fof(f3766, plain, (~ spl66_9 | ~ spl66_44 | ~ spl66_63 | spl66_165), inference(avatar_contradiction_clause, [], [f3765])).
fof(f3765, plain, ($false | (~ spl66_9 | ~ spl66_44 | ~ spl66_63 | spl66_165)), inference(subsumption_resolution, [], [f3764, f867])).
fof(f3764, plain, (~ (op(e0, e0) = e2) | (~ spl66_9 | ~ spl66_44 | spl66_165)), inference(backward_demodulation, [], [f3378, f638])).
fof(f638, plain, ((e0 = op(e3, e1)) | ~ spl66_9), inference(avatar_component_clause, [], [f636])).
fof(f3378, plain, (~ (e2 = op(op(e3, e1), op(e3, e1))) | (~ spl66_44 | spl66_165)), inference(forward_demodulation, [], [f1606, f786])).
fof(f3759, plain, (~ spl66_11 | ~ spl66_15), inference(avatar_split_clause, [], [f3757, f661, f644])).
fof(f3757, plain, (~ (e2 = op(e3, e1)) | ~ spl66_15), inference(backward_demodulation, [], [f277, f663])).
fof(f3752, plain, (~ spl66_21 | ~ spl66_23 | spl66_160), inference(avatar_split_clause, [], [f3747, f1579, f695, f687])).
fof(f3747, plain, (~ (e0 = op(e2, e2)) | (~ spl66_23 | spl66_160)), inference(backward_demodulation, [], [f1581, f697])).
fof(f697, plain, ((e2 = op(e2, e2)) | ~ spl66_23), inference(avatar_component_clause, [], [f695])).
fof(f3748, plain, (~ spl66_19 | ~ spl66_23), inference(avatar_split_clause, [], [f3742, f695, f678])).
fof(f3742, plain, (~ (e2 = op(e2, e3)) | ~ spl66_23), inference(backward_demodulation, [], [f276, f697])).
fof(f3733, plain, (~ spl66_30 | ~ spl66_32), inference(avatar_contradiction_clause, [], [f3732])).
fof(f3732, plain, ($false | (~ spl66_30 | ~ spl66_32)), inference(subsumption_resolution, [], [f3731, f287])).
fof(f287, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f3731, plain, ((e1 = e3) | (~ spl66_30 | ~ spl66_32)), inference(backward_demodulation, [], [f735, f727])).
fof(f3728, plain, (~ spl66_21 | ~ spl66_2 | ~ spl66_35 | spl66_153), inference(avatar_split_clause, [], [f3724, f1546, f746, f606, f687])).
fof(f1546, plain, (spl66_153 <=> (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_153])])).
fof(f3724, plain, (~ (e0 = op(e2, e2)) | (~ spl66_2 | ~ spl66_35 | spl66_153)), inference(backward_demodulation, [], [f3547, f748])).
fof(f3547, plain, (~ (e0 = op(op(e1, e3), op(e1, e3))) | (~ spl66_2 | spl66_153)), inference(backward_demodulation, [], [f1548, f608])).
fof(f1548, plain, (~ (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3))) | spl66_153), inference(avatar_component_clause, [], [f1546])).
fof(f3720, plain, (~ spl66_13 | ~ spl66_45), inference(avatar_split_clause, [], [f3715, f789, f653])).
fof(f3715, plain, (~ (e0 = op(e3, e0)) | ~ spl66_45), inference(backward_demodulation, [], [f239, f791])).
fof(f3701, plain, (~ spl66_30 | ~ spl66_63 | spl66_164), inference(avatar_split_clause, [], [f3694, f1599, f865, f725])).
fof(f1599, plain, (spl66_164 <=> (e1 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_164])])).
fof(f3694, plain, (~ (e1 = op(e2, e0)) | (~ spl66_63 | spl66_164)), inference(backward_demodulation, [], [f1601, f867])).
fof(f1601, plain, (~ (e1 = op(op(e0, e0), e0)) | spl66_164), inference(avatar_component_clause, [], [f1599])).
fof(f3696, plain, (~ spl66_15 | ~ spl66_63), inference(avatar_split_clause, [], [f3688, f865, f661])).
fof(f3688, plain, (~ (e2 = op(e3, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f237, f867])).
fof(f237, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3665, plain, (~ spl66_11 | ~ spl66_21 | ~ spl66_44 | spl66_147), inference(avatar_contradiction_clause, [], [f3664])).
fof(f3664, plain, ($false | (~ spl66_11 | ~ spl66_21 | ~ spl66_44 | spl66_147)), inference(subsumption_resolution, [], [f3662, f689])).
fof(f3662, plain, (~ (e0 = op(e2, e2)) | (~ spl66_11 | ~ spl66_44 | spl66_147)), inference(backward_demodulation, [], [f3379, f646])).
fof(f3379, plain, (~ (e0 = op(op(e3, e1), op(e3, e1))) | (~ spl66_44 | spl66_147)), inference(forward_demodulation, [], [f1520, f786])).
fof(f3627, plain, (spl66_37 | ~ spl66_47 | ~ spl66_74), inference(avatar_contradiction_clause, [], [f3626])).
fof(f3626, plain, ($false | (spl66_37 | ~ spl66_47 | ~ spl66_74)), inference(subsumption_resolution, [], [f3623, f756])).
fof(f756, plain, (~ (e0 = op(e1, e2)) | spl66_37), inference(avatar_component_clause, [], [f755])).
fof(f3623, plain, ((e0 = op(e1, e2)) | (~ spl66_47 | ~ spl66_74)), inference(backward_demodulation, [], [f957, f799])).
fof(f957, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl66_74), inference(avatar_component_clause, [], [f955])).
fof(f955, plain, (spl66_74 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_74])])).
fof(f3625, plain, (~ spl66_35 | ~ spl66_47), inference(avatar_split_clause, [], [f3621, f797, f746])).
fof(f3621, plain, (~ (e2 = op(e1, e3)) | ~ spl66_47), inference(backward_demodulation, [], [f267, f799])).
fof(f3613, plain, (~ spl66_9 | ~ spl66_57), inference(avatar_split_clause, [], [f3607, f840, f636])).
fof(f3607, plain, (~ (e0 = op(e3, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f243, f842])).
fof(f243, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3602, plain, (~ spl66_30 | ~ spl66_62), inference(avatar_split_clause, [], [f3593, f861, f725])).
fof(f3593, plain, (~ (e1 = op(e2, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f236, f863])).
fof(f3587, plain, (~ spl66_9 | ~ spl66_44 | spl66_166), inference(avatar_split_clause, [], [f3586, f1609, f784, f636])).
fof(f3586, plain, (~ (e0 = op(e3, e1)) | (~ spl66_44 | spl66_166)), inference(forward_demodulation, [], [f1611, f786])).
fof(f3577, plain, (~ spl66_51 | ~ spl66_55), inference(avatar_split_clause, [], [f3576, f831, f814])).
fof(f3576, plain, (~ (e2 = op(e0, e3)) | ~ spl66_55), inference(forward_demodulation, [], [f264, f833])).
fof(f3571, plain, (~ spl66_33 | ~ spl66_2 | spl66_162), inference(avatar_split_clause, [], [f3551, f1589, f606, f738])).
fof(f3551, plain, (~ (e0 = op(e1, e3)) | (~ spl66_2 | spl66_162)), inference(backward_demodulation, [], [f1591, f608])).
fof(f1591, plain, (~ (e0 = op(op(e3, e3), e3)) | spl66_162), inference(avatar_component_clause, [], [f1589])).
fof(f3557, plain, (~ spl66_2 | ~ spl66_35 | spl66_154), inference(avatar_contradiction_clause, [], [f3556])).
fof(f3556, plain, ($false | (~ spl66_2 | ~ spl66_35 | spl66_154)), inference(subsumption_resolution, [], [f3548, f748])).
fof(f3548, plain, (~ (e2 = op(e1, e3)) | (~ spl66_2 | spl66_154)), inference(backward_demodulation, [], [f1552, f608])).
fof(f1552, plain, (~ (e2 = op(op(e3, e3), e3)) | spl66_154), inference(avatar_component_clause, [], [f1550])).
fof(f1550, plain, (spl66_154 <=> (e2 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_154])])).
fof(f3542, plain, (~ spl66_5 | ~ spl66_9), inference(avatar_split_clause, [], [f3534, f636, f619])).
fof(f3534, plain, (~ (e0 = op(e3, e2)) | ~ spl66_9), inference(backward_demodulation, [], [f280, f638])).
fof(f3532, plain, (~ spl66_3 | ~ spl66_15), inference(avatar_split_clause, [], [f3529, f661, f610])).
fof(f3529, plain, (~ (e2 = op(e3, e3)) | ~ spl66_15), inference(backward_demodulation, [], [f279, f663])).
fof(f3510, plain, (~ spl66_22 | ~ spl66_30), inference(avatar_split_clause, [], [f3507, f725, f691])).
fof(f3507, plain, (~ (e1 = op(e2, e2)) | ~ spl66_30), inference(backward_demodulation, [], [f272, f727])).
fof(f3504, plain, (~ spl66_3 | ~ spl66_35), inference(avatar_split_clause, [], [f3502, f746, f610])).
fof(f3502, plain, (~ (e2 = op(e3, e3)) | ~ spl66_35), inference(backward_demodulation, [], [f257, f748])).
fof(f257, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3501, plain, (~ spl66_34 | ~ spl66_38), inference(avatar_split_clause, [], [f3496, f759, f742])).
fof(f759, plain, (spl66_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_38])])).
fof(f3496, plain, (~ (e1 = op(e1, e3)) | ~ spl66_38), inference(backward_demodulation, [], [f270, f761])).
fof(f761, plain, ((e1 = op(e1, e2)) | ~ spl66_38), inference(avatar_component_clause, [], [f759])).
fof(f270, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3500, plain, (~ spl66_22 | ~ spl66_38), inference(avatar_split_clause, [], [f3494, f759, f691])).
fof(f3494, plain, (~ (e1 = op(e2, e2)) | ~ spl66_38), inference(backward_demodulation, [], [f250, f761])).
fof(f3493, plain, (~ spl66_39 | ~ spl66_55), inference(avatar_split_clause, [], [f3488, f831, f763])).
fof(f3488, plain, (~ (e2 = op(e1, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f247, f833])).
fof(f247, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3487, plain, (~ spl66_10 | ~ spl66_58), inference(avatar_split_clause, [], [f3482, f844, f640])).
fof(f3482, plain, (~ (e1 = op(e3, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f243, f846])).
fof(f3478, plain, (~ spl66_16 | ~ spl66_64), inference(avatar_split_clause, [], [f3472, f869, f665])).
fof(f3472, plain, (~ (e3 = op(e3, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f237, f871])).
fof(f3466, plain, (~ spl66_22 | spl66_156 | ~ spl66_158), inference(avatar_split_clause, [], [f3465, f1569, f1560, f691])).
fof(f1569, plain, (spl66_158 <=> (e2 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_158])])).
fof(f3465, plain, (~ (e1 = op(e2, e2)) | (spl66_156 | ~ spl66_158)), inference(forward_demodulation, [], [f1562, f1570])).
fof(f1570, plain, ((e2 = op(op(e0, e0), e0)) | ~ spl66_158), inference(avatar_component_clause, [], [f1569])).
fof(f3458, plain, (spl66_64 | ~ spl66_49 | ~ spl66_76), inference(avatar_split_clause, [], [f3395, f967, f806, f869])).
fof(f967, plain, (spl66_76 <=> (e3 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_76])])).
fof(f3395, plain, ((op(e0, e0) = e3) | (~ spl66_49 | ~ spl66_76)), inference(backward_demodulation, [], [f969, f808])).
fof(f969, plain, ((e3 = op(e0, op(e0, e3))) | ~ spl66_76), inference(avatar_component_clause, [], [f967])).
fof(f3449, plain, (spl66_34 | ~ spl66_44 | ~ spl66_73), inference(avatar_split_clause, [], [f3311, f950, f784, f742])).
fof(f950, plain, (spl66_73 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_73])])).
fof(f3311, plain, ((e1 = op(e1, e3)) | (~ spl66_44 | ~ spl66_73)), inference(backward_demodulation, [], [f952, f786])).
fof(f952, plain, ((e1 = op(e1, op(e1, e1))) | ~ spl66_73), inference(avatar_component_clause, [], [f950])).
fof(f3408, plain, (~ spl66_17 | ~ spl66_25), inference(avatar_split_clause, [], [f3402, f704, f670])).
fof(f3402, plain, (~ (e0 = op(e2, e3)) | ~ spl66_25), inference(backward_demodulation, [], [f275, f706])).
fof(f3407, plain, (~ spl66_9 | ~ spl66_25), inference(avatar_split_clause, [], [f3400, f704, f636])).
fof(f3400, plain, (~ (e0 = op(e3, e1)) | ~ spl66_25), inference(backward_demodulation, [], [f246, f706])).
fof(f3396, plain, (~ spl66_17 | ~ spl66_49), inference(avatar_split_clause, [], [f3391, f806, f670])).
fof(f3391, plain, (~ (e0 = op(e2, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f254, f808])).
fof(f3390, plain, (~ spl66_53 | ~ spl66_56), inference(avatar_contradiction_clause, [], [f3389])).
fof(f3389, plain, ($false | (~ spl66_53 | ~ spl66_56)), inference(subsumption_resolution, [], [f3388, f285])).
fof(f285, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f3388, plain, ((e0 = e3) | (~ spl66_53 | ~ spl66_56)), inference(forward_demodulation, [], [f837, f825])).
fof(f3374, plain, (~ spl66_17 | ~ spl66_3 | spl66_162), inference(avatar_split_clause, [], [f3364, f1589, f610, f670])).
fof(f3364, plain, (~ (e0 = op(e2, e3)) | (~ spl66_3 | spl66_162)), inference(forward_demodulation, [], [f1591, f612])).
fof(f3354, plain, (~ spl66_9 | spl66_14 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f3353])).
fof(f3353, plain, ($false | (~ spl66_9 | spl66_14 | ~ spl66_145)), inference(subsumption_resolution, [], [f3352, f658])).
fof(f658, plain, (~ (e1 = op(e3, e0)) | spl66_14), inference(avatar_component_clause, [], [f657])).
fof(f3352, plain, ((e1 = op(e3, e0)) | (~ spl66_9 | ~ spl66_145)), inference(backward_demodulation, [], [f1506, f638])).
fof(f3347, plain, (~ spl66_3 | ~ spl66_17 | ~ spl66_62 | spl66_161), inference(avatar_contradiction_clause, [], [f3346])).
fof(f3346, plain, ($false | (~ spl66_3 | ~ spl66_17 | ~ spl66_62 | spl66_161)), inference(subsumption_resolution, [], [f3341, f863])).
fof(f3341, plain, (~ (op(e0, e0) = e1) | (~ spl66_3 | ~ spl66_17 | spl66_161)), inference(backward_demodulation, [], [f3280, f672])).
fof(f672, plain, ((e0 = op(e2, e3)) | ~ spl66_17), inference(avatar_component_clause, [], [f670])).
fof(f3280, plain, (~ (e1 = op(op(e2, e3), op(e2, e3))) | (~ spl66_3 | spl66_161)), inference(forward_demodulation, [], [f1586, f612])).
fof(f1586, plain, (~ (e1 = op(op(op(e3, e3), e3), op(op(e3, e3), e3))) | spl66_161), inference(avatar_component_clause, [], [f1584])).
fof(f1584, plain, (spl66_161 <=> (e1 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_161])])).
fof(f3307, plain, (~ spl66_52 | spl66_76), inference(avatar_contradiction_clause, [], [f3306])).
fof(f3306, plain, ($false | (~ spl66_52 | spl66_76)), inference(subsumption_resolution, [], [f3304, f820])).
fof(f3304, plain, (~ (e3 = op(e0, e3)) | (~ spl66_52 | spl66_76)), inference(backward_demodulation, [], [f968, f820])).
fof(f968, plain, (~ (e3 = op(e0, op(e0, e3))) | spl66_76), inference(avatar_component_clause, [], [f967])).
fof(f3301, plain, (~ spl66_51 | ~ spl66_59), inference(avatar_split_clause, [], [f3297, f848, f814])).
fof(f3297, plain, (~ (e2 = op(e0, e3)) | ~ spl66_59), inference(backward_demodulation, [], [f263, f850])).
fof(f3265, plain, (spl66_44 | ~ spl66_34 | ~ spl66_71), inference(avatar_split_clause, [], [f3214, f940, f742, f784])).
fof(f940, plain, (spl66_71 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_71])])).
fof(f3214, plain, ((e3 = op(e1, e1)) | (~ spl66_34 | ~ spl66_71)), inference(backward_demodulation, [], [f942, f744])).
fof(f942, plain, ((e3 = op(e1, op(e1, e3))) | ~ spl66_71), inference(avatar_component_clause, [], [f940])).
fof(f3255, plain, (~ spl66_1 | ~ spl66_3), inference(avatar_contradiction_clause, [], [f3254])).
fof(f3254, plain, ($false | (~ spl66_1 | ~ spl66_3)), inference(subsumption_resolution, [], [f3253, f284])).
fof(f3253, plain, ((e0 = e2) | (~ spl66_1 | ~ spl66_3)), inference(backward_demodulation, [], [f612, f604])).
fof(f3234, plain, (~ spl66_20 | spl66_66), inference(avatar_contradiction_clause, [], [f3233])).
fof(f3233, plain, ($false | (~ spl66_20 | spl66_66)), inference(subsumption_resolution, [], [f3232, f684])).
fof(f3232, plain, (~ (e3 = op(e2, e3)) | (~ spl66_20 | spl66_66)), inference(backward_demodulation, [], [f914, f684])).
fof(f914, plain, (~ (e3 = op(e2, op(e2, e3))) | spl66_66), inference(avatar_component_clause, [], [f913])).
fof(f913, plain, (spl66_66 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_66])])).
fof(f3215, plain, (~ spl66_2 | ~ spl66_34), inference(avatar_split_clause, [], [f3211, f742, f606])).
fof(f3211, plain, (~ (e1 = op(e3, e3)) | ~ spl66_34), inference(backward_demodulation, [], [f257, f744])).
fof(f3208, plain, (spl66_35 | ~ spl66_40 | ~ spl66_72), inference(avatar_contradiction_clause, [], [f3207])).
fof(f3207, plain, ($false | (spl66_35 | ~ spl66_40 | ~ spl66_72)), inference(subsumption_resolution, [], [f3204, f747])).
fof(f747, plain, (~ (e2 = op(e1, e3)) | spl66_35), inference(avatar_component_clause, [], [f746])).
fof(f3204, plain, ((e2 = op(e1, e3)) | (~ spl66_40 | ~ spl66_72)), inference(backward_demodulation, [], [f947, f769])).
fof(f947, plain, ((e2 = op(e1, op(e1, e2))) | ~ spl66_72), inference(avatar_component_clause, [], [f945])).
fof(f945, plain, (spl66_72 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_72])])).
fof(f3206, plain, (~ spl66_36 | ~ spl66_40), inference(avatar_split_clause, [], [f3203, f767, f750])).
fof(f3203, plain, (~ (e3 = op(e1, e3)) | ~ spl66_40), inference(backward_demodulation, [], [f270, f769])).
fof(f3192, plain, (~ spl66_37 | ~ spl66_53), inference(avatar_split_clause, [], [f3188, f823, f755])).
fof(f3188, plain, (~ (e0 = op(e1, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f247, f825])).
fof(f3184, plain, (~ spl66_56 | ~ spl66_60), inference(avatar_split_clause, [], [f3181, f852, f835])).
fof(f3181, plain, (~ (e3 = op(e0, e2)) | ~ spl66_60), inference(backward_demodulation, [], [f262, f854])).
fof(f3166, plain, (spl66_38 | ~ spl66_43 | ~ spl66_73), inference(avatar_contradiction_clause, [], [f3165])).
fof(f3165, plain, ($false | (spl66_38 | ~ spl66_43 | ~ spl66_73)), inference(subsumption_resolution, [], [f3164, f760])).
fof(f760, plain, (~ (e1 = op(e1, e2)) | spl66_38), inference(avatar_component_clause, [], [f759])).
fof(f3164, plain, ((e1 = op(e1, e2)) | (~ spl66_43 | ~ spl66_73)), inference(forward_demodulation, [], [f952, f782])).
fof(f3158, plain, (~ spl66_22 | ~ spl66_154 | spl66_161), inference(avatar_contradiction_clause, [], [f3157])).
fof(f3157, plain, ($false | (~ spl66_22 | ~ spl66_154 | spl66_161)), inference(subsumption_resolution, [], [f3154, f693])).
fof(f3154, plain, (~ (e1 = op(e2, e2)) | (~ spl66_154 | spl66_161)), inference(backward_demodulation, [], [f1586, f1551])).
fof(f1551, plain, ((e2 = op(op(e3, e3), e3)) | ~ spl66_154), inference(avatar_component_clause, [], [f1550])).
fof(f3137, plain, (~ spl66_7 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f3136])).
fof(f3136, plain, ($false | (~ spl66_7 | ~ spl66_8)), inference(subsumption_resolution, [], [f3135, f288])).
fof(f288, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f3135, plain, ((e2 = e3) | (~ spl66_7 | ~ spl66_8)), inference(forward_demodulation, [], [f633, f629])).
fof(f3119, plain, (~ spl66_24 | ~ spl66_28), inference(avatar_split_clause, [], [f3114, f716, f699])).
fof(f3114, plain, (~ (e3 = op(e2, e2)) | ~ spl66_28), inference(backward_demodulation, [], [f274, f718])).
fof(f3112, plain, (~ spl66_31 | ~ spl66_32), inference(avatar_contradiction_clause, [], [f3111])).
fof(f3111, plain, ($false | (~ spl66_31 | ~ spl66_32)), inference(subsumption_resolution, [], [f3110, f288])).
fof(f3110, plain, ((e2 = e3) | (~ spl66_31 | ~ spl66_32)), inference(forward_demodulation, [], [f735, f731])).
fof(f3106, plain, (~ spl66_57 | ~ spl66_58), inference(avatar_contradiction_clause, [], [f3105])).
fof(f3105, plain, ($false | (~ spl66_57 | ~ spl66_58)), inference(subsumption_resolution, [], [f3104, f283])).
fof(f283, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f3104, plain, ((e0 = e1) | (~ spl66_57 | ~ spl66_58)), inference(backward_demodulation, [], [f846, f842])).
fof(f3102, plain, (~ spl66_54 | ~ spl66_58), inference(avatar_split_clause, [], [f3097, f844, f827])).
fof(f3097, plain, (~ (e1 = op(e0, e2)) | ~ spl66_58), inference(backward_demodulation, [], [f262, f846])).
fof(f3085, plain, (~ spl66_59 | ~ spl66_63), inference(avatar_split_clause, [], [f3076, f865, f848])).
fof(f3076, plain, (~ (e2 = op(e0, e1)) | ~ spl66_63), inference(backward_demodulation, [], [f259, f867])).
fof(f3049, plain, (spl66_56 | ~ spl66_51 | ~ spl66_76), inference(avatar_split_clause, [], [f2975, f967, f814, f835])).
fof(f2975, plain, ((e3 = op(e0, e2)) | (~ spl66_51 | ~ spl66_76)), inference(backward_demodulation, [], [f969, f816])).
fof(f3043, plain, (~ spl66_28 | ~ spl66_43 | spl66_148), inference(avatar_split_clause, [], [f2928, f1522, f780, f716])).
fof(f2928, plain, (~ (e3 = op(e2, e1)) | (~ spl66_43 | spl66_148)), inference(forward_demodulation, [], [f1524, f782])).
fof(f1524, plain, (~ (e3 = op(op(e1, e1), e1)) | spl66_148), inference(avatar_component_clause, [], [f1522])).
fof(f3033, plain, (spl66_12 | ~ spl66_2 | ~ spl66_143), inference(avatar_split_clause, [], [f3027, f1494, f606, f648])).
fof(f3027, plain, ((e3 = op(e3, e1)) | (~ spl66_2 | ~ spl66_143)), inference(backward_demodulation, [], [f1496, f608])).
fof(f3016, plain, (~ spl66_1 | ~ spl66_13), inference(avatar_split_clause, [], [f3008, f653, f602])).
fof(f3008, plain, (~ (e0 = op(e3, e3)) | ~ spl66_13), inference(backward_demodulation, [], [f279, f655])).
fof(f3006, plain, (~ spl66_17 | ~ spl66_31 | ~ spl66_66), inference(avatar_contradiction_clause, [], [f3005])).
fof(f3005, plain, ($false | (~ spl66_17 | ~ spl66_31 | ~ spl66_66)), inference(subsumption_resolution, [], [f3004, f288])).
fof(f3004, plain, ((e2 = e3) | (~ spl66_17 | ~ spl66_31 | ~ spl66_66)), inference(forward_demodulation, [], [f3002, f731])).
fof(f3002, plain, ((e3 = op(e2, e0)) | (~ spl66_17 | ~ spl66_66)), inference(backward_demodulation, [], [f915, f672])).
fof(f915, plain, ((e3 = op(e2, op(e2, e3))) | ~ spl66_66), inference(avatar_component_clause, [], [f913])).
fof(f3003, plain, (~ spl66_1 | ~ spl66_17), inference(avatar_split_clause, [], [f2998, f670, f602])).
fof(f2998, plain, (~ (e0 = op(e3, e3)) | ~ spl66_17), inference(backward_demodulation, [], [f258, f672])).
fof(f2994, plain, (~ spl66_19 | ~ spl66_31), inference(avatar_split_clause, [], [f2990, f729, f678])).
fof(f2990, plain, (~ (e2 = op(e2, e3)) | ~ spl66_31), inference(backward_demodulation, [], [f273, f731])).
fof(f2993, plain, (~ spl66_15 | ~ spl66_31), inference(avatar_split_clause, [], [f2989, f729, f661])).
fof(f2989, plain, (~ (e2 = op(e3, e0)) | ~ spl66_31), inference(backward_demodulation, [], [f240, f731])).
fof(f2976, plain, (~ spl66_19 | ~ spl66_51), inference(avatar_split_clause, [], [f2972, f814, f678])).
fof(f2972, plain, (~ (e2 = op(e2, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f254, f816])).
fof(f2970, plain, (~ spl66_53 | ~ spl66_54), inference(avatar_contradiction_clause, [], [f2969])).
fof(f2969, plain, ($false | (~ spl66_53 | ~ spl66_54)), inference(subsumption_resolution, [], [f2968, f283])).
fof(f2968, plain, ((e0 = e1) | (~ spl66_53 | ~ spl66_54)), inference(forward_demodulation, [], [f829, f825])).
fof(f2965, plain, (~ spl66_53 | ~ spl66_59 | ~ spl66_78), inference(avatar_contradiction_clause, [], [f2964])).
fof(f2964, plain, ($false | (~ spl66_53 | ~ spl66_59 | ~ spl66_78)), inference(subsumption_resolution, [], [f2963, f283])).
fof(f2963, plain, ((e0 = e1) | (~ spl66_53 | ~ spl66_59 | ~ spl66_78)), inference(forward_demodulation, [], [f2961, f825])).
fof(f2961, plain, ((e1 = op(e0, e2)) | (~ spl66_59 | ~ spl66_78)), inference(backward_demodulation, [], [f979, f850])).
fof(f2959, plain, (~ spl66_15 | ~ spl66_64 | spl66_158), inference(avatar_split_clause, [], [f2953, f1569, f869, f661])).
fof(f2953, plain, (~ (e2 = op(e3, e0)) | (~ spl66_64 | spl66_158)), inference(backward_demodulation, [], [f1571, f871])).
fof(f1571, plain, (~ (e2 = op(op(e0, e0), e0)) | spl66_158), inference(avatar_component_clause, [], [f1569])).
fof(f2958, plain, (~ spl66_52 | ~ spl66_64), inference(avatar_split_clause, [], [f2950, f869, f818])).
fof(f2950, plain, (~ (e3 = op(e0, e3)) | ~ spl66_64), inference(backward_demodulation, [], [f261, f871])).
fof(f2945, plain, (~ spl66_19 | ~ spl66_24 | spl66_67), inference(avatar_split_clause, [], [f2944, f918, f699, f678])).
fof(f2944, plain, (~ (e2 = op(e2, e3)) | (~ spl66_24 | spl66_67)), inference(forward_demodulation, [], [f919, f701])).
fof(f919, plain, (~ (e2 = op(e2, op(e2, e2))) | spl66_67), inference(avatar_component_clause, [], [f918])).
fof(f2943, plain, (~ spl66_26 | spl66_68), inference(avatar_contradiction_clause, [], [f2942])).
fof(f2942, plain, ($false | (~ spl66_26 | spl66_68)), inference(subsumption_resolution, [], [f2941, f710])).
fof(f2941, plain, (~ (e1 = op(e2, e1)) | (~ spl66_26 | spl66_68)), inference(forward_demodulation, [], [f924, f710])).
fof(f924, plain, (~ (e1 = op(e2, op(e2, e1))) | spl66_68), inference(avatar_component_clause, [], [f923])).
fof(f923, plain, (spl66_68 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_68])])).
fof(f2936, plain, (~ spl66_6 | ~ spl66_24 | spl66_152), inference(avatar_split_clause, [], [f2935, f1541, f699, f623])).
fof(f2935, plain, (~ (e1 = op(e3, e2)) | (~ spl66_24 | spl66_152)), inference(forward_demodulation, [], [f1543, f701])).
fof(f2913, plain, (spl66_63 | ~ spl66_53 | ~ spl66_77), inference(avatar_split_clause, [], [f2912, f972, f823, f865])).
fof(f972, plain, (spl66_77 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_77])])).
fof(f2912, plain, ((op(e0, e0) = e2) | (~ spl66_53 | ~ spl66_77)), inference(forward_demodulation, [], [f974, f825])).
fof(f974, plain, ((e2 = op(e0, op(e0, e2))) | ~ spl66_77), inference(avatar_component_clause, [], [f972])).
fof(f2905, plain, (~ spl66_35 | ~ spl66_43), inference(avatar_split_clause, [], [f2904, f780, f746])).
fof(f2904, plain, (~ (e2 = op(e1, e3)) | ~ spl66_43), inference(forward_demodulation, [], [f269, f782])).
fof(f2890, plain, (~ spl66_6 | spl66_11 | ~ spl66_144), inference(avatar_contradiction_clause, [], [f2889])).
fof(f2889, plain, ($false | (~ spl66_6 | spl66_11 | ~ spl66_144)), inference(subsumption_resolution, [], [f2884, f645])).
fof(f645, plain, (~ (e2 = op(e3, e1)) | spl66_11), inference(avatar_component_clause, [], [f644])).
fof(f2884, plain, ((e2 = op(e3, e1)) | (~ spl66_6 | ~ spl66_144)), inference(backward_demodulation, [], [f1501, f625])).
fof(f625, plain, ((e1 = op(e3, e2)) | ~ spl66_6), inference(avatar_component_clause, [], [f623])).
fof(f2888, plain, (~ spl66_2 | ~ spl66_6), inference(avatar_split_clause, [], [f2881, f623, f606])).
fof(f2881, plain, (~ (e1 = op(e3, e3)) | ~ spl66_6), inference(backward_demodulation, [], [f282, f625])).
fof(f2863, plain, (~ spl66_5 | ~ spl66_37), inference(avatar_split_clause, [], [f2858, f755, f619])).
fof(f2858, plain, (~ (e0 = op(e3, e2)) | ~ spl66_37), inference(backward_demodulation, [], [f251, f757])).
fof(f757, plain, ((e0 = op(e1, e2)) | ~ spl66_37), inference(avatar_component_clause, [], [f755])).
fof(f2834, plain, (~ spl66_53 | ~ spl66_55), inference(avatar_contradiction_clause, [], [f2833])).
fof(f2833, plain, ($false | (~ spl66_53 | ~ spl66_55)), inference(subsumption_resolution, [], [f2832, f284])).
fof(f2832, plain, ((e0 = e2) | (~ spl66_53 | ~ spl66_55)), inference(backward_demodulation, [], [f833, f825])).
fof(f2814, plain, (~ spl66_47 | ~ spl66_62 | spl66_158), inference(avatar_split_clause, [], [f2808, f1569, f861, f797])).
fof(f2808, plain, (~ (e2 = op(e1, e0)) | (~ spl66_62 | spl66_158)), inference(backward_demodulation, [], [f1571, f863])).
fof(f2796, plain, (~ spl66_5 | ~ spl66_24 | spl66_160), inference(avatar_split_clause, [], [f2795, f1579, f699, f619])).
fof(f2795, plain, (~ (e0 = op(e3, e2)) | (~ spl66_24 | spl66_160)), inference(forward_demodulation, [], [f1581, f701])).
fof(f2765, plain, (~ spl66_10 | spl66_145), inference(avatar_contradiction_clause, [], [f2764])).
fof(f2764, plain, ($false | (~ spl66_10 | spl66_145)), inference(subsumption_resolution, [], [f2763, f642])).
fof(f2763, plain, (~ (e1 = op(e3, e1)) | (~ spl66_10 | spl66_145)), inference(backward_demodulation, [], [f1505, f642])).
fof(f1505, plain, (~ (e1 = op(e3, op(e3, e1))) | spl66_145), inference(avatar_component_clause, [], [f1504])).
fof(f2746, plain, (~ spl66_41 | ~ spl66_44), inference(avatar_contradiction_clause, [], [f2745])).
fof(f2745, plain, ($false | (~ spl66_41 | ~ spl66_44)), inference(subsumption_resolution, [], [f2744, f285])).
fof(f2744, plain, ((e0 = e3) | (~ spl66_41 | ~ spl66_44)), inference(forward_demodulation, [], [f786, f774])).
fof(f2743, plain, (~ spl66_49 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f2742])).
fof(f2742, plain, ($false | (~ spl66_49 | ~ spl66_52)), inference(subsumption_resolution, [], [f2741, f285])).
fof(f2741, plain, ((e0 = e3) | (~ spl66_49 | ~ spl66_52)), inference(forward_demodulation, [], [f820, f808])).
fof(f2739, plain, (~ spl66_24 | ~ spl66_41 | ~ spl66_59 | spl66_169), inference(avatar_contradiction_clause, [], [f2738])).
fof(f2738, plain, ($false | (~ spl66_24 | ~ spl66_41 | ~ spl66_59 | spl66_169)), inference(subsumption_resolution, [], [f2734, f701])).
fof(f2734, plain, (~ (e3 = op(e2, e2)) | (~ spl66_41 | ~ spl66_59 | spl66_169)), inference(backward_demodulation, [], [f2711, f850])).
fof(f2711, plain, (~ (e3 = op(op(e0, e1), op(e0, e1))) | (~ spl66_41 | spl66_169)), inference(forward_demodulation, [], [f1628, f774])).
fof(f2718, plain, (spl66_15 | ~ spl66_5 | ~ spl66_144), inference(avatar_split_clause, [], [f2717, f1499, f619, f661])).
fof(f2717, plain, ((e2 = op(e3, e0)) | (~ spl66_5 | ~ spl66_144)), inference(forward_demodulation, [], [f1501, f621])).
fof(f2704, plain, (~ spl66_60 | ~ spl66_41 | spl66_148), inference(avatar_split_clause, [], [f2613, f1522, f772, f852])).
fof(f2613, plain, (~ (e3 = op(e0, e1)) | (~ spl66_41 | spl66_148)), inference(backward_demodulation, [], [f1524, f774])).
fof(f2698, plain, (spl66_59 | ~ spl66_54 | ~ spl66_77), inference(avatar_split_clause, [], [f2697, f972, f827, f848])).
fof(f2697, plain, ((e2 = op(e0, e1)) | (~ spl66_54 | ~ spl66_77)), inference(forward_demodulation, [], [f974, f829])).
fof(f2686, plain, (~ spl66_39 | ~ spl66_47), inference(avatar_split_clause, [], [f2685, f797, f763])).
fof(f2685, plain, (~ (e2 = op(e1, e2)) | ~ spl66_47), inference(forward_demodulation, [], [f266, f799])).
fof(f266, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f2674, plain, (~ spl66_25 | ~ spl66_41), inference(avatar_split_clause, [], [f2606, f772, f704])).
fof(f2606, plain, (~ (e0 = op(e2, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f244, f774])).
fof(f2647, plain, (~ spl66_15 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f2646])).
fof(f2646, plain, ($false | (~ spl66_15 | ~ spl66_16)), inference(subsumption_resolution, [], [f2645, f288])).
fof(f2645, plain, ((e2 = e3) | (~ spl66_15 | ~ spl66_16)), inference(forward_demodulation, [], [f667, f663])).
fof(f2627, plain, (~ spl66_34 | ~ spl66_36), inference(avatar_contradiction_clause, [], [f2626])).
fof(f2626, plain, ($false | (~ spl66_34 | ~ spl66_36)), inference(subsumption_resolution, [], [f2625, f287])).
fof(f2625, plain, ((e1 = e3) | (~ spl66_34 | ~ spl66_36)), inference(backward_demodulation, [], [f752, f744])).
fof(f752, plain, ((e3 = op(e1, e3)) | ~ spl66_36), inference(avatar_component_clause, [], [f750])).
fof(f2623, plain, (~ spl66_20 | ~ spl66_36), inference(avatar_split_clause, [], [f2621, f750, f682])).
fof(f2621, plain, (~ (e3 = op(e2, e3)) | ~ spl66_36), inference(backward_demodulation, [], [f256, f752])).
fof(f2589, plain, (~ spl66_54 | ~ spl66_55), inference(avatar_contradiction_clause, [], [f2588])).
fof(f2588, plain, ($false | (~ spl66_54 | ~ spl66_55)), inference(subsumption_resolution, [], [f2587, f286])).
fof(f286, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f2587, plain, ((e1 = e2) | (~ spl66_54 | ~ spl66_55)), inference(backward_demodulation, [], [f833, f829])).
fof(f2586, plain, (~ spl66_50 | ~ spl66_60 | spl66_78), inference(avatar_split_clause, [], [f2584, f977, f852, f810])).
fof(f2584, plain, (~ (e1 = op(e0, e3)) | (~ spl66_60 | spl66_78)), inference(backward_demodulation, [], [f978, f854])).
fof(f978, plain, (~ (e1 = op(e0, op(e0, e1))) | spl66_78), inference(avatar_component_clause, [], [f977])).
fof(f2585, plain, (~ spl66_12 | ~ spl66_60), inference(avatar_split_clause, [], [f2582, f852, f648])).
fof(f2582, plain, (~ (e3 = op(e3, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f243, f854])).
fof(f2576, plain, (~ spl66_50 | ~ spl66_62), inference(avatar_split_clause, [], [f2569, f861, f810])).
fof(f2569, plain, (~ (e1 = op(e0, e3)) | ~ spl66_62), inference(backward_demodulation, [], [f261, f863])).
fof(f2566, plain, (spl66_5 | ~ spl66_15 | ~ spl66_146), inference(avatar_split_clause, [], [f2565, f1509, f661, f619])).
fof(f2565, plain, ((e0 = op(e3, e2)) | (~ spl66_15 | ~ spl66_146)), inference(forward_demodulation, [], [f1511, f663])).
fof(f2563, plain, (~ spl66_62 | spl66_159 | ~ spl66_160), inference(avatar_split_clause, [], [f2562, f1579, f1574, f861])).
fof(f2562, plain, (~ (op(e0, e0) = e1) | (spl66_159 | ~ spl66_160)), inference(forward_demodulation, [], [f1576, f1580])).
fof(f2561, plain, (~ spl66_24 | ~ spl66_158 | spl66_168), inference(avatar_split_clause, [], [f2560, f1620, f1569, f699])).
fof(f2560, plain, (~ (e3 = op(e2, e2)) | (~ spl66_158 | spl66_168)), inference(forward_demodulation, [], [f1622, f1570])).
fof(f2528, plain, (spl66_22 | ~ spl66_27 | ~ spl66_68), inference(avatar_split_clause, [], [f2525, f923, f712, f691])).
fof(f2525, plain, ((e1 = op(e2, e2)) | (~ spl66_27 | ~ spl66_68)), inference(backward_demodulation, [], [f925, f714])).
fof(f925, plain, ((e1 = op(e2, op(e2, e1))) | ~ spl66_68), inference(avatar_component_clause, [], [f923])).
fof(f2522, plain, (~ spl66_35 | ~ spl66_37 | ~ spl66_71), inference(avatar_contradiction_clause, [], [f2521])).
fof(f2521, plain, ($false | (~ spl66_35 | ~ spl66_37 | ~ spl66_71)), inference(subsumption_resolution, [], [f2520, f285])).
fof(f2520, plain, ((e0 = e3) | (~ spl66_35 | ~ spl66_37 | ~ spl66_71)), inference(forward_demodulation, [], [f2518, f757])).
fof(f2518, plain, ((e3 = op(e1, e2)) | (~ spl66_35 | ~ spl66_71)), inference(backward_demodulation, [], [f942, f748])).
fof(f2507, plain, (~ spl66_42 | spl66_73), inference(avatar_contradiction_clause, [], [f2506])).
fof(f2506, plain, ($false | (~ spl66_42 | spl66_73)), inference(subsumption_resolution, [], [f2500, f778])).
fof(f778, plain, ((e1 = op(e1, e1)) | ~ spl66_42), inference(avatar_component_clause, [], [f776])).
fof(f776, plain, (spl66_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_42])])).
fof(f2500, plain, (~ (e1 = op(e1, e1)) | (~ spl66_42 | spl66_73)), inference(backward_demodulation, [], [f951, f778])).
fof(f951, plain, (~ (e1 = op(e1, op(e1, e1))) | spl66_73), inference(avatar_component_clause, [], [f950])).
fof(f2505, plain, (~ spl66_26 | ~ spl66_42), inference(avatar_split_clause, [], [f2496, f776, f708])).
fof(f2496, plain, (~ (e1 = op(e2, e1)) | ~ spl66_42), inference(backward_demodulation, [], [f244, f778])).
fof(f2491, plain, (~ spl66_43 | ~ spl66_1 | ~ spl66_50 | spl66_167), inference(avatar_split_clause, [], [f2490, f1614, f810, f602, f780])).
fof(f2490, plain, (~ (e2 = op(e1, e1)) | (~ spl66_1 | ~ spl66_50 | spl66_167)), inference(forward_demodulation, [], [f2489, f812])).
fof(f2489, plain, (~ (e2 = op(op(e0, e3), op(e0, e3))) | (~ spl66_1 | spl66_167)), inference(forward_demodulation, [], [f1616, f604])).
fof(f2485, plain, (~ spl66_42 | ~ spl66_46), inference(avatar_split_clause, [], [f2484, f793, f776])).
fof(f2484, plain, (~ (e1 = op(e1, e1)) | ~ spl66_46), inference(forward_demodulation, [], [f265, f795])).
fof(f2467, plain, (~ spl66_1 | ~ spl66_50 | spl66_155), inference(avatar_contradiction_clause, [], [f2466])).
fof(f2466, plain, ($false | (~ spl66_1 | ~ spl66_50 | spl66_155)), inference(subsumption_resolution, [], [f2465, f812])).
fof(f2465, plain, (~ (e1 = op(e0, e3)) | (~ spl66_1 | spl66_155)), inference(forward_demodulation, [], [f1557, f604])).
fof(f2447, plain, (spl66_49 | ~ spl66_64 | ~ spl66_79), inference(avatar_contradiction_clause, [], [f2446])).
fof(f2446, plain, ($false | (spl66_49 | ~ spl66_64 | ~ spl66_79)), inference(subsumption_resolution, [], [f2440, f807])).
fof(f807, plain, (~ (e0 = op(e0, e3)) | spl66_49), inference(avatar_component_clause, [], [f806])).
fof(f2440, plain, ((e0 = op(e0, e3)) | (~ spl66_64 | ~ spl66_79)), inference(backward_demodulation, [], [f984, f871])).
fof(f2415, plain, (~ spl66_14 | ~ spl66_6), inference(avatar_split_clause, [], [f2414, f623, f657])).
fof(f2414, plain, (~ (e1 = op(e3, e0)) | ~ spl66_6), inference(forward_demodulation, [], [f278, f625])).
fof(f2393, plain, (~ spl66_21 | ~ spl66_29), inference(avatar_split_clause, [], [f2390, f721, f687])).
fof(f2390, plain, (~ (e0 = op(e2, e2)) | ~ spl66_29), inference(backward_demodulation, [], [f272, f723])).
fof(f2379, plain, (~ spl66_46 | ~ spl66_48), inference(avatar_contradiction_clause, [], [f2378])).
fof(f2378, plain, ($false | (~ spl66_46 | ~ spl66_48)), inference(subsumption_resolution, [], [f2377, f287])).
fof(f2377, plain, ((e1 = e3) | (~ spl66_46 | ~ spl66_48)), inference(forward_demodulation, [], [f803, f795])).
fof(f2376, plain, (~ spl66_50 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f2375])).
fof(f2375, plain, ($false | (~ spl66_50 | ~ spl66_52)), inference(subsumption_resolution, [], [f2374, f287])).
fof(f2374, plain, ((e1 = e3) | (~ spl66_50 | ~ spl66_52)), inference(backward_demodulation, [], [f820, f812])).
fof(f2357, plain, (spl66_62 | ~ spl66_57 | ~ spl66_78), inference(avatar_split_clause, [], [f2125, f977, f840, f861])).
fof(f2125, plain, ((op(e0, e0) = e1) | (~ spl66_57 | ~ spl66_78)), inference(forward_demodulation, [], [f979, f842])).
fof(f2355, plain, (~ spl66_43 | spl66_163 | ~ spl66_164), inference(avatar_contradiction_clause, [], [f2354])).
fof(f2354, plain, ($false | (~ spl66_43 | spl66_163 | ~ spl66_164)), inference(subsumption_resolution, [], [f2353, f782])).
fof(f2353, plain, (~ (e2 = op(e1, e1)) | (spl66_163 | ~ spl66_164)), inference(forward_demodulation, [], [f1596, f1600])).
fof(f1600, plain, ((e1 = op(op(e0, e0), e0)) | ~ spl66_164), inference(avatar_component_clause, [], [f1599])).
fof(f2341, plain, (~ spl66_21 | ~ spl66_37), inference(avatar_split_clause, [], [f2340, f755, f687])).
fof(f2340, plain, (~ (e0 = op(e2, e2)) | ~ spl66_37), inference(forward_demodulation, [], [f250, f757])).
fof(f2333, plain, (~ spl66_33 | ~ spl66_37), inference(avatar_split_clause, [], [f2332, f755, f738])).
fof(f2332, plain, (~ (e0 = op(e1, e3)) | ~ spl66_37), inference(forward_demodulation, [], [f270, f757])).
fof(f2328, plain, (~ spl66_33 | ~ spl66_1), inference(avatar_split_clause, [], [f2239, f602, f738])).
fof(f2239, plain, (~ (e0 = op(e1, e3)) | ~ spl66_1), inference(forward_demodulation, [], [f257, f604])).
fof(f2320, plain, (spl66_24 | ~ spl66_19 | ~ spl66_66), inference(avatar_split_clause, [], [f2230, f913, f678, f699])).
fof(f2230, plain, ((e3 = op(e2, e2)) | (~ spl66_19 | ~ spl66_66)), inference(forward_demodulation, [], [f915, f680])).
fof(f2312, plain, (~ spl66_13 | spl66_146), inference(avatar_contradiction_clause, [], [f2311])).
fof(f2311, plain, ($false | (~ spl66_13 | spl66_146)), inference(subsumption_resolution, [], [f2310, f655])).
fof(f2310, plain, (~ (e0 = op(e3, e0)) | (~ spl66_13 | spl66_146)), inference(forward_demodulation, [], [f1510, f655])).
fof(f1510, plain, (~ (e0 = op(e3, op(e3, e0))) | spl66_146), inference(avatar_component_clause, [], [f1509])).
fof(f2303, plain, (~ spl66_1 | ~ spl66_2), inference(avatar_contradiction_clause, [], [f2302])).
fof(f2302, plain, ($false | (~ spl66_1 | ~ spl66_2)), inference(subsumption_resolution, [], [f2301, f283])).
fof(f2301, plain, ((e0 = e1) | (~ spl66_1 | ~ spl66_2)), inference(forward_demodulation, [], [f608, f604])).
fof(f2300, plain, (~ spl66_1 | ~ spl66_12 | ~ spl66_145), inference(avatar_contradiction_clause, [], [f2299])).
fof(f2299, plain, ($false | (~ spl66_1 | ~ spl66_12 | ~ spl66_145)), inference(subsumption_resolution, [], [f2298, f283])).
fof(f2298, plain, ((e0 = e1) | (~ spl66_1 | ~ spl66_12 | ~ spl66_145)), inference(forward_demodulation, [], [f2297, f604])).
fof(f2297, plain, ((e1 = op(e3, e3)) | (~ spl66_12 | ~ spl66_145)), inference(backward_demodulation, [], [f1506, f650])).
fof(f2292, plain, (~ spl66_17 | ~ spl66_19), inference(avatar_contradiction_clause, [], [f2291])).
fof(f2291, plain, ($false | (~ spl66_17 | ~ spl66_19)), inference(subsumption_resolution, [], [f2290, f284])).
fof(f2290, plain, ((e0 = e2) | (~ spl66_17 | ~ spl66_19)), inference(backward_demodulation, [], [f680, f672])).
fof(f2280, plain, (~ spl66_51 | ~ spl66_1 | spl66_154), inference(avatar_split_clause, [], [f2279, f1550, f602, f814])).
fof(f2279, plain, (~ (e2 = op(e0, e3)) | (~ spl66_1 | spl66_154)), inference(forward_demodulation, [], [f1552, f604])).
fof(f2270, plain, (~ spl66_38 | ~ spl66_46), inference(avatar_contradiction_clause, [], [f2269])).
fof(f2269, plain, ($false | (~ spl66_38 | ~ spl66_46)), inference(subsumption_resolution, [], [f1921, f795])).
fof(f1921, plain, (~ (e1 = op(e1, e0)) | ~ spl66_38), inference(forward_demodulation, [], [f266, f761])).
fof(f2258, plain, (spl66_41 | ~ spl66_46 | ~ spl66_74), inference(avatar_contradiction_clause, [], [f2257])).
fof(f2257, plain, ($false | (spl66_41 | ~ spl66_46 | ~ spl66_74)), inference(subsumption_resolution, [], [f2256, f773])).
fof(f773, plain, (~ (e0 = op(e1, e1)) | spl66_41), inference(avatar_component_clause, [], [f772])).
fof(f2256, plain, ((e0 = op(e1, e1)) | (~ spl66_46 | ~ spl66_74)), inference(forward_demodulation, [], [f957, f795])).
fof(f2235, plain, (~ spl66_36 | ~ spl66_52), inference(avatar_split_clause, [], [f2234, f818, f750])).
fof(f2234, plain, (~ (e3 = op(e1, e3)) | ~ spl66_52), inference(forward_demodulation, [], [f253, f820])).
fof(f2222, plain, (~ spl66_17 | ~ spl66_21), inference(avatar_split_clause, [], [f2014, f687, f670])).
fof(f2014, plain, (~ (e0 = op(e2, e3)) | ~ spl66_21), inference(forward_demodulation, [], [f276, f689])).
fof(f2221, plain, (spl66_17 | ~ spl66_32 | ~ spl66_69), inference(avatar_split_clause, [], [f2178, f928, f733, f670])).
fof(f2178, plain, ((e0 = op(e2, e3)) | (~ spl66_32 | ~ spl66_69)), inference(backward_demodulation, [], [f930, f735])).
fof(f2183, plain, (~ spl66_19 | ~ spl66_20), inference(avatar_contradiction_clause, [], [f2182])).
fof(f2182, plain, ($false | (~ spl66_19 | ~ spl66_20)), inference(subsumption_resolution, [], [f2181, f288])).
fof(f2181, plain, ((e2 = e3) | (~ spl66_19 | ~ spl66_20)), inference(backward_demodulation, [], [f684, f680])).
fof(f2174, plain, (~ spl66_43 | ~ spl66_44), inference(avatar_contradiction_clause, [], [f2173])).
fof(f2173, plain, ($false | (~ spl66_43 | ~ spl66_44)), inference(subsumption_resolution, [], [f2172, f288])).
fof(f2172, plain, ((e2 = e3) | (~ spl66_43 | ~ spl66_44)), inference(backward_demodulation, [], [f786, f782])).
fof(f2171, plain, (~ spl66_45 | ~ spl66_46), inference(avatar_contradiction_clause, [], [f2170])).
fof(f2170, plain, ($false | (~ spl66_45 | ~ spl66_46)), inference(subsumption_resolution, [], [f2169, f283])).
fof(f2169, plain, ((e0 = e1) | (~ spl66_45 | ~ spl66_46)), inference(forward_demodulation, [], [f795, f791])).
fof(f2163, plain, (spl66_43 | ~ spl66_38 | ~ spl66_72), inference(avatar_split_clause, [], [f2162, f945, f759, f780])).
fof(f2162, plain, ((e2 = op(e1, e1)) | (~ spl66_38 | ~ spl66_72)), inference(forward_demodulation, [], [f947, f761])).
fof(f2124, plain, (~ spl66_50 | ~ spl66_34), inference(avatar_split_clause, [], [f2123, f742, f810])).
fof(f2123, plain, (~ (e1 = op(e0, e3)) | ~ spl66_34), inference(forward_demodulation, [], [f253, f744])).
fof(f2074, plain, (~ spl66_34 | ~ spl66_35), inference(avatar_contradiction_clause, [], [f2073])).
fof(f2073, plain, ($false | (~ spl66_34 | ~ spl66_35)), inference(subsumption_resolution, [], [f2072, f286])).
fof(f2072, plain, ((e1 = e2) | (~ spl66_34 | ~ spl66_35)), inference(backward_demodulation, [], [f748, f744])).
fof(f2070, plain, (~ spl66_38 | ~ spl66_40), inference(avatar_contradiction_clause, [], [f2069])).
fof(f2069, plain, ($false | (~ spl66_38 | ~ spl66_40)), inference(subsumption_resolution, [], [f2068, f287])).
fof(f2068, plain, ((e1 = e3) | (~ spl66_38 | ~ spl66_40)), inference(forward_demodulation, [], [f769, f761])).
fof(f2067, plain, (~ spl66_57 | ~ spl66_59), inference(avatar_contradiction_clause, [], [f2066])).
fof(f2066, plain, ($false | (~ spl66_57 | ~ spl66_59)), inference(subsumption_resolution, [], [f2065, f284])).
fof(f2065, plain, ((e0 = e2) | (~ spl66_57 | ~ spl66_59)), inference(backward_demodulation, [], [f850, f842])).
fof(f2063, plain, (~ spl66_62 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f2062])).
fof(f2062, plain, ($false | (~ spl66_62 | ~ spl66_64)), inference(subsumption_resolution, [], [f2061, f287])).
fof(f2061, plain, ((e1 = e3) | (~ spl66_62 | ~ spl66_64)), inference(backward_demodulation, [], [f871, f863])).
fof(f2056, plain, (spl66_54 | ~ spl66_59 | ~ spl66_78), inference(avatar_contradiction_clause, [], [f2055])).
fof(f2055, plain, ($false | (spl66_54 | ~ spl66_59 | ~ spl66_78)), inference(subsumption_resolution, [], [f2054, f828])).
fof(f828, plain, (~ (e1 = op(e0, e2)) | spl66_54), inference(avatar_component_clause, [], [f827])).
fof(f2054, plain, ((e1 = op(e0, e2)) | (~ spl66_59 | ~ spl66_78)), inference(forward_demodulation, [], [f979, f850])).
fof(f2039, plain, (~ spl66_55 | ~ spl66_59), inference(avatar_contradiction_clause, [], [f2038])).
fof(f2038, plain, ($false | (~ spl66_55 | ~ spl66_59)), inference(subsumption_resolution, [], [f2037, f850])).
fof(f2037, plain, (~ (e2 = op(e0, e1)) | ~ spl66_55), inference(forward_demodulation, [], [f262, f833])).
fof(f1989, plain, (~ spl66_9 | ~ spl66_12), inference(avatar_contradiction_clause, [], [f1988])).
fof(f1988, plain, ($false | (~ spl66_9 | ~ spl66_12)), inference(subsumption_resolution, [], [f1987, f285])).
fof(f1987, plain, ((e0 = e3) | (~ spl66_9 | ~ spl66_12)), inference(forward_demodulation, [], [f650, f638])).
fof(f1986, plain, (~ spl66_21 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f1985])).
fof(f1985, plain, ($false | (~ spl66_21 | ~ spl66_24)), inference(subsumption_resolution, [], [f1983, f285])).
fof(f1983, plain, ((e0 = e3) | (~ spl66_21 | ~ spl66_24)), inference(backward_demodulation, [], [f701, f689])).
fof(f1958, plain, (~ spl66_28 | ~ spl66_44), inference(avatar_split_clause, [], [f1953, f784, f716])).
fof(f1953, plain, (~ (e3 = op(e2, e1)) | ~ spl66_44), inference(backward_demodulation, [], [f244, f786])).
fof(f1944, plain, (~ spl66_41 | spl66_150 | ~ spl66_152), inference(avatar_split_clause, [], [f1943, f1541, f1532, f772])).
fof(f1943, plain, (~ (e0 = op(e1, e1)) | (spl66_150 | ~ spl66_152)), inference(forward_demodulation, [], [f1534, f1542])).
fof(f1542, plain, ((e1 = op(op(e2, e2), e2)) | ~ spl66_152), inference(avatar_component_clause, [], [f1541])).
fof(f1932, plain, (~ spl66_41 | ~ spl66_3 | ~ spl66_18 | spl66_153), inference(avatar_split_clause, [], [f1898, f1546, f674, f610, f772])).
fof(f1898, plain, (~ (e0 = op(e1, e1)) | (~ spl66_3 | ~ spl66_18 | spl66_153)), inference(forward_demodulation, [], [f1892, f676])).
fof(f1892, plain, (~ (e0 = op(op(e2, e3), op(e2, e3))) | (~ spl66_3 | spl66_153)), inference(backward_demodulation, [], [f1548, f612])).
fof(f1927, plain, (~ spl66_6 | ~ spl66_38), inference(avatar_split_clause, [], [f1926, f759, f623])).
fof(f1926, plain, (~ (e1 = op(e3, e2)) | ~ spl66_38), inference(forward_demodulation, [], [f251, f761])).
fof(f1923, plain, (~ spl66_42 | ~ spl66_38), inference(avatar_split_clause, [], [f1922, f759, f776])).
fof(f1922, plain, (~ (e1 = op(e1, e1)) | ~ spl66_38), inference(forward_demodulation, [], [f268, f761])).
fof(f1915, plain, (spl66_28 | ~ spl66_18 | ~ spl66_66), inference(avatar_split_clause, [], [f1792, f913, f674, f716])).
fof(f1792, plain, ((e3 = op(e2, e1)) | (~ spl66_18 | ~ spl66_66)), inference(forward_demodulation, [], [f915, f676])).
fof(f1888, plain, (~ spl66_5 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f1887])).
fof(f1887, plain, ($false | (~ spl66_5 | ~ spl66_8)), inference(subsumption_resolution, [], [f1886, f285])).
fof(f1886, plain, ((e0 = e3) | (~ spl66_5 | ~ spl66_8)), inference(forward_demodulation, [], [f633, f621])).
fof(f1885, plain, (~ spl66_1 | ~ spl66_9), inference(avatar_split_clause, [], [f1884, f636, f602])).
fof(f1884, plain, (~ (e0 = op(e3, e3)) | ~ spl66_9), inference(backward_demodulation, [], [f281, f638])).
fof(f1860, plain, (~ spl66_38 | ~ spl66_39), inference(avatar_contradiction_clause, [], [f1859])).
fof(f1859, plain, ($false | (~ spl66_38 | ~ spl66_39)), inference(subsumption_resolution, [], [f1858, f286])).
fof(f1858, plain, ((e1 = e2) | (~ spl66_38 | ~ spl66_39)), inference(backward_demodulation, [], [f765, f761])).
fof(f1856, plain, (~ spl66_41 | ~ spl66_43), inference(avatar_contradiction_clause, [], [f1855])).
fof(f1855, plain, ($false | (~ spl66_41 | ~ spl66_43)), inference(subsumption_resolution, [], [f1854, f284])).
fof(f1854, plain, ((e0 = e2) | (~ spl66_41 | ~ spl66_43)), inference(forward_demodulation, [], [f782, f774])).
fof(f1834, plain, (spl66_9 | ~ spl66_14 | ~ spl66_146), inference(avatar_split_clause, [], [f1833, f1509, f657, f636])).
fof(f1833, plain, ((e0 = op(e3, e1)) | (~ spl66_14 | ~ spl66_146)), inference(forward_demodulation, [], [f1511, f659])).
fof(f1818, plain, (~ spl66_47 | ~ spl66_31), inference(avatar_split_clause, [], [f1817, f729, f797])).
fof(f1817, plain, (~ (e2 = op(e1, e0)) | ~ spl66_31), inference(forward_demodulation, [], [f238, f731])).
fof(f1810, plain, (~ spl66_48 | ~ spl66_36), inference(avatar_split_clause, [], [f1809, f750, f801])).
fof(f1809, plain, (~ (e3 = op(e1, e0)) | ~ spl66_36), inference(forward_demodulation, [], [f267, f752])).
fof(f1799, plain, (~ spl66_59 | ~ spl66_41 | spl66_149), inference(avatar_split_clause, [], [f1798, f1527, f772, f848])).
fof(f1798, plain, (~ (e2 = op(e0, e1)) | (~ spl66_41 | spl66_149)), inference(forward_demodulation, [], [f1529, f774])).
fof(f1789, plain, (~ spl66_10 | ~ spl66_14), inference(avatar_split_clause, [], [f1788, f657, f640])).
fof(f1788, plain, (~ (e1 = op(e3, e1)) | ~ spl66_14), inference(forward_demodulation, [], [f277, f659])).
fof(f1754, plain, (~ spl66_17 | ~ spl66_18), inference(avatar_contradiction_clause, [], [f1753])).
fof(f1753, plain, ($false | (~ spl66_17 | ~ spl66_18)), inference(subsumption_resolution, [], [f1752, f283])).
fof(f1752, plain, ((e0 = e1) | (~ spl66_17 | ~ spl66_18)), inference(backward_demodulation, [], [f676, f672])).
fof(f1751, plain, (~ spl66_18 | ~ spl66_19), inference(avatar_contradiction_clause, [], [f1750])).
fof(f1750, plain, ($false | (~ spl66_18 | ~ spl66_19)), inference(subsumption_resolution, [], [f1749, f286])).
fof(f1749, plain, ((e1 = e2) | (~ spl66_18 | ~ spl66_19)), inference(backward_demodulation, [], [f680, f676])).
fof(f1735, plain, (spl66_18 | ~ spl66_28 | ~ spl66_68), inference(avatar_split_clause, [], [f1731, f923, f716, f674])).
fof(f1731, plain, ((e1 = op(e2, e3)) | (~ spl66_28 | ~ spl66_68)), inference(backward_demodulation, [], [f925, f718])).
fof(f1734, plain, (~ spl66_20 | ~ spl66_28), inference(avatar_split_clause, [], [f1730, f716, f682])).
fof(f1730, plain, (~ (e3 = op(e2, e3)) | ~ spl66_28), inference(backward_demodulation, [], [f275, f718])).
fof(f1727, plain, (spl66_21 | ~ spl66_31 | ~ spl66_69), inference(avatar_split_clause, [], [f1722, f928, f729, f687])).
fof(f1722, plain, ((e0 = op(e2, e2)) | (~ spl66_31 | ~ spl66_69)), inference(backward_demodulation, [], [f930, f731])).
fof(f1717, plain, (~ spl66_4 | ~ spl66_36), inference(avatar_split_clause, [], [f1715, f750, f614])).
fof(f1715, plain, (~ (e3 = op(e3, e3)) | ~ spl66_36), inference(backward_demodulation, [], [f257, f752])).
fof(f1713, plain, (~ spl66_35 | ~ spl66_39), inference(avatar_split_clause, [], [f1710, f763, f746])).
fof(f1710, plain, (~ (e2 = op(e1, e3)) | ~ spl66_39), inference(backward_demodulation, [], [f270, f765])).
fof(f1711, plain, (~ spl66_23 | ~ spl66_39), inference(avatar_split_clause, [], [f1708, f763, f695])).
fof(f1708, plain, (~ (e2 = op(e2, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f250, f765])).
fof(f1707, plain, (~ spl66_41 | ~ spl66_42), inference(avatar_contradiction_clause, [], [f1706])).
fof(f1706, plain, ($false | (~ spl66_41 | ~ spl66_42)), inference(subsumption_resolution, [], [f1705, f283])).
fof(f1705, plain, ((e0 = e1) | (~ spl66_41 | ~ spl66_42)), inference(backward_demodulation, [], [f778, f774])).
fof(f1701, plain, (~ spl66_34 | ~ spl66_42), inference(avatar_split_clause, [], [f1693, f776, f742])).
fof(f1693, plain, (~ (e1 = op(e1, e3)) | ~ spl66_42), inference(backward_demodulation, [], [f269, f778])).
fof(f1699, plain, (~ spl66_10 | ~ spl66_42), inference(avatar_split_clause, [], [f1691, f776, f640])).
fof(f1691, plain, (~ (e1 = op(e3, e1)) | ~ spl66_42), inference(backward_demodulation, [], [f245, f778])).
fof(f1672, plain, (~ spl66_50 | ~ spl66_54), inference(avatar_split_clause, [], [f1668, f827, f810])).
fof(f1668, plain, (~ (e1 = op(e0, e3)) | ~ spl66_54), inference(backward_demodulation, [], [f264, f829])).
fof(f1671, plain, (~ spl66_6 | ~ spl66_54), inference(avatar_split_clause, [], [f1667, f827, f623])).
fof(f1667, plain, (~ (e1 = op(e3, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f249, f829])).
fof(f1669, plain, (~ spl66_38 | ~ spl66_54), inference(avatar_split_clause, [], [f1665, f827, f759])).
fof(f1665, plain, (~ (e1 = op(e1, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f247, f829])).
fof(f1660, plain, (~ spl66_44 | ~ spl66_60), inference(avatar_split_clause, [], [f1655, f852, f784])).
fof(f1655, plain, (~ (e3 = op(e1, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f241, f854])).
fof(f1654, plain, (~ spl66_63 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f1653])).
fof(f1653, plain, ($false | (~ spl66_63 | ~ spl66_64)), inference(subsumption_resolution, [], [f1652, f288])).
fof(f1652, plain, ((e2 = e3) | (~ spl66_63 | ~ spl66_64)), inference(backward_demodulation, [], [f871, f867])).
fof(f1651, plain, (~ spl66_14 | ~ spl66_64 | spl66_164), inference(avatar_split_clause, [], [f1644, f1599, f869, f657])).
fof(f1644, plain, (~ (e1 = op(e3, e0)) | (~ spl66_64 | spl66_164)), inference(backward_demodulation, [], [f1601, f871])).
fof(f1648, plain, (~ spl66_60 | ~ spl66_64), inference(avatar_split_clause, [], [f1640, f869, f852])).
fof(f1640, plain, (~ (e3 = op(e0, e1)) | ~ spl66_64), inference(backward_demodulation, [], [f259, f871])).
fof(f1646, plain, (~ spl66_32 | ~ spl66_64), inference(avatar_split_clause, [], [f1638, f869, f733])).
fof(f1638, plain, (~ (e3 = op(e2, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f236, f871])).
fof(f1636, plain, (~ spl66_170 | ~ spl66_22 | ~ spl66_160), inference(avatar_split_clause, [], [f600, f1579, f691, f1632])).
fof(f600, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e0 = op(op(e2, e2), e2)) & (e1 = op(e2, e2)) & (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax29)).
fof(f1635, plain, (~ spl66_170 | ~ spl66_21 | ~ spl66_152), inference(avatar_split_clause, [], [f599, f1541, f687, f1632])).
fof(f599, plain, (~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e1 = op(op(e2, e2), e2)) & (e0 = op(e2, e2)) & (e3 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax28)).
fof(f1630, plain, (~ spl66_169 | ~ spl66_43 | ~ spl66_166), inference(avatar_split_clause, [], [f598, f1609, f780, f1626])).
fof(f598, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e0 = op(op(e1, e1), e1)) & (e2 = op(e1, e1)) & (e3 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax27)).
fof(f1629, plain, (~ spl66_169 | ~ spl66_41 | ~ spl66_149), inference(avatar_split_clause, [], [f597, f1527, f772, f1626])).
fof(f597, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(op(e1, e1), e1)) & (e0 = op(e1, e1)) & (e3 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax26)).
fof(f1624, plain, (~ spl66_168 | ~ spl66_63 | ~ spl66_164), inference(avatar_split_clause, [], [f596, f1599, f865, f1620])).
fof(f596, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e2) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f49])).
fof(f49, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e2) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e1 = op(op(e0, e0), e0)) & (op(e0, e0) = e2) & (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax25)).
fof(f1623, plain, (~ spl66_168 | ~ spl66_62 | ~ spl66_158), inference(avatar_split_clause, [], [f595, f1569, f861, f1620])).
fof(f595, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e2 = op(op(e0, e0), e0)) & (op(e0, e0) = e1) & (e3 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax24)).
fof(f1618, plain, (~ spl66_167 | ~ spl66_2 | ~ spl66_162), inference(avatar_split_clause, [], [f594, f1589, f606, f1614])).
fof(f594, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(op(e3, e3), e3)) & (e1 = op(e3, e3)) & (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax23)).
fof(f1617, plain, (~ spl66_167 | ~ spl66_1 | ~ spl66_155), inference(avatar_split_clause, [], [f593, f1555, f602, f1614])).
fof(f593, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(op(e3, e3), e3)) & (e0 = op(e3, e3)) & (e2 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax22)).
fof(f1612, plain, (~ spl66_165 | ~ spl66_44 | ~ spl66_166), inference(avatar_split_clause, [], [f592, f1609, f784, f1604])).
fof(f592, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(op(e1, e1), e1)) & (e3 = op(e1, e1)) & (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax21)).
fof(f1607, plain, (~ spl66_165 | ~ spl66_41 | ~ spl66_148), inference(avatar_split_clause, [], [f591, f1522, f772, f1604])).
fof(f591, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e3 = op(op(e1, e1), e1)) & (e0 = op(e1, e1)) & (e2 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax20)).
fof(f1602, plain, (~ spl66_163 | ~ spl66_64 | ~ spl66_164), inference(avatar_split_clause, [], [f590, f1599, f869, f1594])).
fof(f590, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(op(e0, e0), e0)) & (op(e0, e0) = e3) & (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax19)).
fof(f1597, plain, (~ spl66_163 | ~ spl66_62 | ~ spl66_157), inference(avatar_split_clause, [], [f589, f1564, f861, f1594])).
fof(f589, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(op(e0, e0), e0)) & (op(e0, e0) = e1) & (e2 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax18)).
fof(f1592, plain, (~ spl66_161 | ~ spl66_3 | ~ spl66_162), inference(avatar_split_clause, [], [f588, f1589, f610, f1584])).
fof(f588, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(e3, e3)) | ~ (e1 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(op(e3, e3), e3)) & (e2 = op(e3, e3)) & (e1 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax17)).
fof(f1587, plain, (~ spl66_161 | ~ spl66_1 | ~ spl66_154), inference(avatar_split_clause, [], [f587, f1550, f602, f1584])).
fof(f587, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e2 = op(op(e3, e3), e3)) & (e0 = op(e3, e3)) & (e1 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax16)).
fof(f1582, plain, (~ spl66_159 | ~ spl66_24 | ~ spl66_160), inference(avatar_split_clause, [], [f586, f1579, f699, f1574])).
fof(f586, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(op(e2, e2), e2)) & (e3 = op(e2, e2)) & (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax15)).
fof(f1577, plain, (~ spl66_159 | ~ spl66_21 | ~ spl66_151), inference(avatar_split_clause, [], [f585, f1536, f687, f1574])).
fof(f585, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(op(e2, e2), e2)) & (e0 = op(e2, e2)) & (e1 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax14)).
fof(f1572, plain, (~ spl66_156 | ~ spl66_64 | ~ spl66_158), inference(avatar_split_clause, [], [f584, f1569, f869, f1560])).
fof(f584, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e2 = op(op(e0, e0), e0)) & (op(e0, e0) = e3) & (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax13)).
fof(f1567, plain, (~ spl66_156 | ~ spl66_63 | ~ spl66_157), inference(avatar_split_clause, [], [f583, f1564, f865, f1560])).
fof(f583, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(op(e0, e0), e0)) & (op(e0, e0) = e2) & (e1 = op(op(op(e0, e0), e0), op(op(e0, e0), e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax12)).
fof(f1558, plain, (~ spl66_153 | ~ spl66_3 | ~ spl66_155), inference(avatar_split_clause, [], [f582, f1555, f610, f1546])).
fof(f582, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e1 = op(op(e3, e3), e3)) & (e2 = op(e3, e3)) & (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax11)).
fof(f1553, plain, (~ spl66_153 | ~ spl66_2 | ~ spl66_154), inference(avatar_split_clause, [], [f581, f1550, f606, f1546])).
fof(f581, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(op(e3, e3), e3)) & (e1 = op(e3, e3)) & (e0 = op(op(op(e3, e3), e3), op(op(e3, e3), e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax10)).
fof(f1544, plain, (~ spl66_150 | ~ spl66_24 | ~ spl66_152), inference(avatar_split_clause, [], [f580, f1541, f699, f1532])).
fof(f580, plain, (~ (e1 = op(op(e2, e2), e2)) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e1 = op(op(e2, e2), e2)) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e1 = op(op(e2, e2), e2)) & (e3 = op(e2, e2)) & (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax9)).
fof(f1539, plain, (~ spl66_150 | ~ spl66_22 | ~ spl66_151), inference(avatar_split_clause, [], [f579, f1536, f691, f1532])).
fof(f579, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(op(e2, e2), e2)) & (e1 = op(e2, e2)) & (e0 = op(op(op(e2, e2), e2), op(op(e2, e2), e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax8)).
fof(f1530, plain, (~ spl66_147 | ~ spl66_44 | ~ spl66_149), inference(avatar_split_clause, [], [f578, f1527, f784, f1518])).
fof(f578, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e2 = op(op(e1, e1), e1)) & (e3 = op(e1, e1)) & (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax7)).
fof(f1525, plain, (~ spl66_147 | ~ spl66_43 | ~ spl66_148), inference(avatar_split_clause, [], [f577, f1522, f780, f1518])).
fof(f577, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e3 = op(op(e1, e1), e1)) & (e2 = op(e1, e1)) & (e0 = op(op(op(e1, e1), e1), op(op(e1, e1), e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax6)).
fof(f1516, plain, (spl66_142 | spl66_141 | spl66_140 | spl66_139 | spl66_138 | spl66_137 | spl66_136 | spl66_135 | spl66_134 | spl66_133 | spl66_132 | spl66_131 | spl66_130 | spl66_129 | spl66_128 | spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | spl66_4), inference(avatar_split_clause, [], [f565, f614, f987, f995, f1003, f1011, f1019, f1027, f1035, f1043, f1051, f1059, f1067, f1075, f1083, f1091, f1099, f1107, f1115, f1123, f1131, f1139, f1147, f1155, f1163, f1171, f1179, f1187, f1195, f1203, f1211, f1219, f1227, f1235, f1243, f1251, f1259, f1267, f1275, f1283, f1291, f1299, f1307, f1315, f1323, f1331, f1339, f1347, f1355, f1363, f1371, f1379, f1387, f1395, f1403, f1411, f1419, f1427, f1435, f1443, f1451, f1459, f1467, f1475, f1483])).
fof(f1483, plain, (spl66_142 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl66_142])])).
fof(f1475, plain, (spl66_141 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl66_141])])).
fof(f1467, plain, (spl66_140 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl66_140])])).
fof(f1459, plain, (spl66_139 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl66_139])])).
fof(f1451, plain, (spl66_138 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl66_138])])).
fof(f1443, plain, (spl66_137 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl66_137])])).
fof(f1435, plain, (spl66_136 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl66_136])])).
fof(f1427, plain, (spl66_135 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl66_135])])).
fof(f1419, plain, (spl66_134 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl66_134])])).
fof(f1411, plain, (spl66_133 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl66_133])])).
fof(f1403, plain, (spl66_132 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl66_132])])).
fof(f1395, plain, (spl66_131 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl66_131])])).
fof(f1387, plain, (spl66_130 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl66_130])])).
fof(f1379, plain, (spl66_129 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl66_129])])).
fof(f1371, plain, (spl66_128 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl66_128])])).
fof(f1363, plain, (spl66_127 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl66_127])])).
fof(f1355, plain, (spl66_126 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl66_126])])).
fof(f1347, plain, (spl66_125 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl66_125])])).
fof(f1339, plain, (spl66_124 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl66_124])])).
fof(f1331, plain, (spl66_123 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl66_123])])).
fof(f1323, plain, (spl66_122 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl66_122])])).
fof(f1315, plain, (spl66_121 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl66_121])])).
fof(f1307, plain, (spl66_120 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl66_120])])).
fof(f1299, plain, (spl66_119 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl66_119])])).
fof(f1291, plain, (spl66_118 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl66_118])])).
fof(f1283, plain, (spl66_117 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl66_117])])).
fof(f1275, plain, (spl66_116 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl66_116])])).
fof(f1267, plain, (spl66_115 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl66_115])])).
fof(f1259, plain, (spl66_114 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl66_114])])).
fof(f1251, plain, (spl66_113 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl66_113])])).
fof(f1243, plain, (spl66_112 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl66_112])])).
fof(f1235, plain, (spl66_111 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl66_111])])).
fof(f1227, plain, (spl66_110 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl66_110])])).
fof(f1219, plain, (spl66_109 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl66_109])])).
fof(f1211, plain, (spl66_108 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl66_108])])).
fof(f1203, plain, (spl66_107 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl66_107])])).
fof(f1195, plain, (spl66_106 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl66_106])])).
fof(f1187, plain, (spl66_105 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl66_105])])).
fof(f1179, plain, (spl66_104 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl66_104])])).
fof(f1171, plain, (spl66_103 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl66_103])])).
fof(f1163, plain, (spl66_102 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl66_102])])).
fof(f1155, plain, (spl66_101 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl66_101])])).
fof(f1147, plain, (spl66_100 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl66_100])])).
fof(f1139, plain, (spl66_99 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl66_99])])).
fof(f1131, plain, (spl66_98 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl66_98])])).
fof(f1123, plain, (spl66_97 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl66_97])])).
fof(f1115, plain, (spl66_96 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl66_96])])).
fof(f1107, plain, (spl66_95 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl66_95])])).
fof(f1099, plain, (spl66_94 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl66_94])])).
fof(f1091, plain, (spl66_93 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl66_93])])).
fof(f1083, plain, (spl66_92 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl66_92])])).
fof(f1075, plain, (spl66_91 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl66_91])])).
fof(f1067, plain, (spl66_90 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl66_90])])).
fof(f1059, plain, (spl66_89 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl66_89])])).
fof(f1051, plain, (spl66_88 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl66_88])])).
fof(f1043, plain, (spl66_87 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl66_87])])).
fof(f1035, plain, (spl66_86 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl66_86])])).
fof(f1027, plain, (spl66_85 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl66_85])])).
fof(f1019, plain, (spl66_84 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl66_84])])).
fof(f1011, plain, (spl66_83 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl66_83])])).
fof(f1003, plain, (spl66_82 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl66_82])])).
fof(f995, plain, (spl66_81 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl66_81])])).
fof(f987, plain, (spl66_80 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl66_80])])).
fof(f565, plain, ((e3 = op(e3, e3)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f120, plain, (((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0)))) | sP65 | sP64 | sP63) & ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f84, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP30), inference(usedef, [], [e84])).
fof(e84, plain, (sP30 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f85, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(usedef, [], [e85])).
fof(e85, plain, (sP31 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f86, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP32), inference(usedef, [], [e86])).
fof(e86, plain, (sP32 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f87, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP33), inference(usedef, [], [e87])).
fof(e87, plain, (sP33 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f88, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP34), inference(usedef, [], [e88])).
fof(e88, plain, (sP34 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f89, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(usedef, [], [e89])).
fof(e89, plain, (sP35 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f90, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP36), inference(usedef, [], [e90])).
fof(e90, plain, (sP36 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f91, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP37), inference(usedef, [], [e91])).
fof(e91, plain, (sP37 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f92, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP38), inference(usedef, [], [e92])).
fof(e92, plain, (sP38 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f93, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(usedef, [], [e93])).
fof(e93, plain, (sP39 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f94, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP40), inference(usedef, [], [e94])).
fof(e94, plain, (sP40 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f95, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP41), inference(usedef, [], [e95])).
fof(e95, plain, (sP41 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f96, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP42), inference(usedef, [], [e96])).
fof(e96, plain, (sP42 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f97, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(usedef, [], [e97])).
fof(e97, plain, (sP43 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f98, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP44), inference(usedef, [], [e98])).
fof(e98, plain, (sP44 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f99, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP45), inference(usedef, [], [e99])).
fof(e99, plain, (sP45 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f100, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP46), inference(usedef, [], [e100])).
fof(e100, plain, (sP46 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f101, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(usedef, [], [e101])).
fof(e101, plain, (sP47 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f102, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP48), inference(usedef, [], [e102])).
fof(e102, plain, (sP48 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f103, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP49), inference(usedef, [], [e103])).
fof(e103, plain, (sP49 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f104, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP50), inference(usedef, [], [e104])).
fof(e104, plain, (sP50 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f105, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(usedef, [], [e105])).
fof(e105, plain, (sP51 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f106, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP52), inference(usedef, [], [e106])).
fof(e106, plain, (sP52 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f107, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP53), inference(usedef, [], [e107])).
fof(e107, plain, (sP53 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f108, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP54), inference(usedef, [], [e108])).
fof(e108, plain, (sP54 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f109, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(usedef, [], [e109])).
fof(e109, plain, (sP55 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f110, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP56), inference(usedef, [], [e110])).
fof(e110, plain, (sP56 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f111, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP57), inference(usedef, [], [e111])).
fof(e111, plain, (sP57 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f112, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP58), inference(usedef, [], [e112])).
fof(e112, plain, (sP58 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f113, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(usedef, [], [e113])).
fof(e113, plain, (sP59 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f114, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP60), inference(usedef, [], [e114])).
fof(e114, plain, (sP60 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f115, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP61), inference(usedef, [], [e115])).
fof(e115, plain, (sP61 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f116, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP62), inference(usedef, [], [e116])).
fof(e116, plain, (sP62 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f117, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0)))) | ~ sP63), inference(usedef, [], [e117])).
fof(e117, plain, (sP63 <=> (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f118, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0)))) | ~ sP64), inference(usedef, [], [e118])).
fof(e118, plain, (sP64 <=> (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f119, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0)))) | ~ sP65), inference(usedef, [], [e119])).
fof(e119, plain, (sP65 <=> (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f5, plain, (((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0)))) | (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0)))) | (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0)))) | (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))))) & ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax5)).
fof(f1515, plain, (spl66_142 | spl66_141 | spl66_140 | spl66_139 | spl66_138 | spl66_137 | spl66_136 | spl66_135 | spl66_134 | spl66_133 | spl66_132 | spl66_131 | spl66_130 | spl66_129 | spl66_128 | spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | ~ spl66_4), inference(avatar_split_clause, [], [f566, f614, f987, f995, f1003, f1011, f1019, f1027, f1035, f1043, f1051, f1059, f1067, f1075, f1083, f1091, f1099, f1107, f1115, f1123, f1131, f1139, f1147, f1155, f1163, f1171, f1179, f1187, f1195, f1203, f1211, f1219, f1227, f1235, f1243, f1251, f1259, f1267, f1275, f1283, f1291, f1299, f1307, f1315, f1323, f1331, f1339, f1347, f1355, f1363, f1371, f1379, f1387, f1395, f1403, f1411, f1419, f1427, f1435, f1443, f1451, f1459, f1467, f1475, f1483])).
fof(f566, plain, (~ (e3 = op(e3, e3)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1512, plain, (spl66_75 | spl66_70 | spl66_65 | spl66_146), inference(avatar_split_clause, [], [f569, f1509, f906, f933, f960])).
fof(f960, plain, (spl66_75 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl66_75])])).
fof(f933, plain, (spl66_70 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl66_70])])).
fof(f906, plain, (spl66_65 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl66_65])])).
fof(f569, plain, ((e0 = op(e3, op(e3, e0))) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1507, plain, (spl66_75 | spl66_70 | spl66_65 | spl66_145), inference(avatar_split_clause, [], [f570, f1504, f906, f933, f960])).
fof(f570, plain, ((e1 = op(e3, op(e3, e1))) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1502, plain, (spl66_75 | spl66_70 | spl66_65 | spl66_144), inference(avatar_split_clause, [], [f571, f1499, f906, f933, f960])).
fof(f571, plain, ((e2 = op(e3, op(e3, e2))) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1497, plain, (spl66_75 | spl66_70 | spl66_65 | spl66_143), inference(avatar_split_clause, [], [f572, f1494, f906, f933, f960])).
fof(f572, plain, ((e3 = op(e3, op(e3, e3))) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1492, plain, (spl66_75 | spl66_70 | spl66_65 | ~ spl66_49 | spl66_64), inference(avatar_split_clause, [], [f573, f869, f806, f906, f933, f960])).
fof(f573, plain, ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3)) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1491, plain, (spl66_75 | spl66_70 | spl66_65 | ~ spl66_34 | spl66_44), inference(avatar_split_clause, [], [f574, f784, f742, f906, f933, f960])).
fof(f574, plain, ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3)) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1490, plain, (spl66_75 | spl66_70 | spl66_65 | ~ spl66_19 | spl66_24), inference(avatar_split_clause, [], [f575, f699, f678, f906, f933, f960])).
fof(f575, plain, ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3)) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1489, plain, (~ spl66_142 | spl66_61), inference(avatar_split_clause, [], [f561, f857, f1483])).
fof(f561, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1488, plain, (~ spl66_142 | ~ spl66_61), inference(avatar_split_clause, [], [f562, f857, f1483])).
fof(f562, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1481, plain, (~ spl66_141 | spl66_61), inference(avatar_split_clause, [], [f557, f857, f1475])).
fof(f557, plain, ((e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1480, plain, (~ spl66_141 | ~ spl66_61), inference(avatar_split_clause, [], [f558, f857, f1475])).
fof(f558, plain, (~ (e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1479, plain, (~ spl66_141 | spl66_41), inference(avatar_split_clause, [], [f559, f772, f1475])).
fof(f559, plain, ((e0 = op(e1, e1)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1473, plain, (~ spl66_140 | spl66_61), inference(avatar_split_clause, [], [f553, f857, f1467])).
fof(f553, plain, ((e0 = op(e0, e0)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1472, plain, (~ spl66_140 | ~ spl66_61), inference(avatar_split_clause, [], [f554, f857, f1467])).
fof(f554, plain, (~ (e0 = op(e0, e0)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1465, plain, (~ spl66_139 | spl66_61), inference(avatar_split_clause, [], [f549, f857, f1459])).
fof(f549, plain, ((e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1464, plain, (~ spl66_139 | ~ spl66_61), inference(avatar_split_clause, [], [f550, f857, f1459])).
fof(f550, plain, (~ (e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f1456, plain, (~ spl66_138 | ~ spl66_62), inference(avatar_split_clause, [], [f546, f861, f1451])).
fof(f546, plain, (~ (op(e0, e0) = e1) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1455, plain, (~ spl66_138 | spl66_62), inference(avatar_split_clause, [], [f547, f861, f1451])).
fof(f547, plain, ((op(e0, e0) = e1) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f1447, plain, (~ spl66_137 | spl66_42), inference(avatar_split_clause, [], [f543, f776, f1443])).
fof(f543, plain, ((e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1446, plain, (~ spl66_137 | ~ spl66_42), inference(avatar_split_clause, [], [f544, f776, f1443])).
fof(f544, plain, (~ (e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f1441, plain, (~ spl66_136 | spl66_57), inference(avatar_split_clause, [], [f537, f840, f1435])).
fof(f537, plain, ((e0 = op(e0, e1)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1440, plain, (~ spl66_136 | ~ spl66_62), inference(avatar_split_clause, [], [f538, f861, f1435])).
fof(f538, plain, (~ (op(e0, e0) = e1) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1439, plain, (~ spl66_136 | spl66_22), inference(avatar_split_clause, [], [f539, f691, f1435])).
fof(f539, plain, ((e1 = op(e2, e2)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1438, plain, (~ spl66_136 | ~ spl66_27), inference(avatar_split_clause, [], [f540, f712, f1435])).
fof(f540, plain, (~ (e2 = op(e2, e1)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1433, plain, (~ spl66_135 | spl66_57), inference(avatar_split_clause, [], [f533, f840, f1427])).
fof(f533, plain, ((e0 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1432, plain, (~ spl66_135 | ~ spl66_62), inference(avatar_split_clause, [], [f534, f861, f1427])).
fof(f534, plain, (~ (op(e0, e0) = e1) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1431, plain, (~ spl66_135 | spl66_2), inference(avatar_split_clause, [], [f535, f606, f1427])).
fof(f535, plain, ((e1 = op(e3, e3)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1430, plain, (~ spl66_135 | ~ spl66_12), inference(avatar_split_clause, [], [f536, f648, f1427])).
fof(f536, plain, (~ (e3 = op(e3, e1)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1425, plain, (~ spl66_134 | spl66_53), inference(avatar_split_clause, [], [f529, f823, f1419])).
fof(f529, plain, ((e0 = op(e0, e2)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1424, plain, (~ spl66_134 | ~ spl66_63), inference(avatar_split_clause, [], [f530, f865, f1419])).
fof(f530, plain, (~ (op(e0, e0) = e2) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f1423, plain, (~ spl66_134 | spl66_63), inference(avatar_split_clause, [], [f531, f865, f1419])).
fof(f531, plain, ((op(e0, e0) = e2) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f1417, plain, (~ spl66_133 | spl66_53), inference(avatar_split_clause, [], [f525, f823, f1411])).
fof(f525, plain, ((e0 = op(e0, e2)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1416, plain, (~ spl66_133 | ~ spl66_63), inference(avatar_split_clause, [], [f526, f865, f1411])).
fof(f526, plain, (~ (op(e0, e0) = e2) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1415, plain, (~ spl66_133 | spl66_43), inference(avatar_split_clause, [], [f527, f780, f1411])).
fof(f527, plain, ((e2 = op(e1, e1)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1414, plain, (~ spl66_133 | ~ spl66_38), inference(avatar_split_clause, [], [f528, f759, f1411])).
fof(f528, plain, (~ (e1 = op(e1, e2)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1409, plain, (~ spl66_132 | spl66_53), inference(avatar_split_clause, [], [f521, f823, f1403])).
fof(f521, plain, ((e0 = op(e0, e2)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f1407, plain, (~ spl66_132 | spl66_23), inference(avatar_split_clause, [], [f523, f695, f1403])).
fof(f523, plain, ((e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1406, plain, (~ spl66_132 | ~ spl66_23), inference(avatar_split_clause, [], [f524, f695, f1403])).
fof(f524, plain, (~ (e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1401, plain, (~ spl66_131 | spl66_53), inference(avatar_split_clause, [], [f517, f823, f1395])).
fof(f517, plain, ((e0 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f1400, plain, (~ spl66_131 | ~ spl66_63), inference(avatar_split_clause, [], [f518, f865, f1395])).
fof(f518, plain, (~ (op(e0, e0) = e2) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1399, plain, (~ spl66_131 | spl66_3), inference(avatar_split_clause, [], [f519, f610, f1395])).
fof(f519, plain, ((e2 = op(e3, e3)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1398, plain, (~ spl66_131 | ~ spl66_8), inference(avatar_split_clause, [], [f520, f631, f1395])).
fof(f520, plain, (~ (e3 = op(e3, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1392, plain, (~ spl66_130 | ~ spl66_64), inference(avatar_split_clause, [], [f514, f869, f1387])).
fof(f514, plain, (~ (op(e0, e0) = e3) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f1391, plain, (~ spl66_130 | spl66_64), inference(avatar_split_clause, [], [f515, f869, f1387])).
fof(f515, plain, ((op(e0, e0) = e3) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1385, plain, (~ spl66_129 | spl66_49), inference(avatar_split_clause, [], [f509, f806, f1379])).
fof(f509, plain, ((e0 = op(e0, e3)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f1384, plain, (~ spl66_129 | ~ spl66_64), inference(avatar_split_clause, [], [f510, f869, f1379])).
fof(f510, plain, (~ (op(e0, e0) = e3) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1383, plain, (~ spl66_129 | spl66_44), inference(avatar_split_clause, [], [f511, f784, f1379])).
fof(f511, plain, ((e3 = op(e1, e1)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1382, plain, (~ spl66_129 | ~ spl66_34), inference(avatar_split_clause, [], [f512, f742, f1379])).
fof(f512, plain, (~ (e1 = op(e1, e3)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1377, plain, (~ spl66_128 | spl66_49), inference(avatar_split_clause, [], [f505, f806, f1371])).
fof(f505, plain, ((e0 = op(e0, e3)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f1376, plain, (~ spl66_128 | ~ spl66_64), inference(avatar_split_clause, [], [f506, f869, f1371])).
fof(f506, plain, (~ (op(e0, e0) = e3) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1375, plain, (~ spl66_128 | spl66_24), inference(avatar_split_clause, [], [f507, f699, f1371])).
fof(f507, plain, ((e3 = op(e2, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1374, plain, (~ spl66_128 | ~ spl66_19), inference(avatar_split_clause, [], [f508, f678, f1371])).
fof(f508, plain, (~ (e2 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1367, plain, (~ spl66_127 | spl66_4), inference(avatar_split_clause, [], [f503, f614, f1363])).
fof(f503, plain, ((e3 = op(e3, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f1366, plain, (~ spl66_127 | ~ spl66_4), inference(avatar_split_clause, [], [f504, f614, f1363])).
fof(f504, plain, (~ (e3 = op(e3, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1361, plain, (~ spl66_126 | spl66_46), inference(avatar_split_clause, [], [f497, f793, f1355])).
fof(f497, plain, ((e1 = op(e1, e0)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f1359, plain, (~ spl66_126 | spl66_61), inference(avatar_split_clause, [], [f499, f857, f1355])).
fof(f499, plain, ((e0 = op(e0, e0)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1358, plain, (~ spl66_126 | ~ spl66_61), inference(avatar_split_clause, [], [f500, f857, f1355])).
fof(f500, plain, (~ (e0 = op(e0, e0)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1353, plain, (~ spl66_125 | spl66_46), inference(avatar_split_clause, [], [f493, f793, f1347])).
fof(f493, plain, ((e1 = op(e1, e0)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f1352, plain, (~ spl66_125 | ~ spl66_41), inference(avatar_split_clause, [], [f494, f772, f1347])).
fof(f494, plain, (~ (e0 = op(e1, e1)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1351, plain, (~ spl66_125 | spl66_41), inference(avatar_split_clause, [], [f495, f772, f1347])).
fof(f495, plain, ((e0 = op(e1, e1)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1345, plain, (~ spl66_124 | spl66_46), inference(avatar_split_clause, [], [f489, f793, f1339])).
fof(f489, plain, ((e1 = op(e1, e0)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f1344, plain, (~ spl66_124 | ~ spl66_41), inference(avatar_split_clause, [], [f490, f772, f1339])).
fof(f490, plain, (~ (e0 = op(e1, e1)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1343, plain, (~ spl66_124 | spl66_21), inference(avatar_split_clause, [], [f491, f687, f1339])).
fof(f491, plain, ((e0 = op(e2, e2)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1342, plain, (~ spl66_124 | ~ spl66_31), inference(avatar_split_clause, [], [f492, f729, f1339])).
fof(f492, plain, (~ (e2 = op(e2, e0)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1337, plain, (~ spl66_123 | spl66_46), inference(avatar_split_clause, [], [f485, f793, f1331])).
fof(f485, plain, ((e1 = op(e1, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f1336, plain, (~ spl66_123 | ~ spl66_41), inference(avatar_split_clause, [], [f486, f772, f1331])).
fof(f486, plain, (~ (e0 = op(e1, e1)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1335, plain, (~ spl66_123 | spl66_1), inference(avatar_split_clause, [], [f487, f602, f1331])).
fof(f487, plain, ((e0 = op(e3, e3)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1334, plain, (~ spl66_123 | ~ spl66_16), inference(avatar_split_clause, [], [f488, f665, f1331])).
fof(f488, plain, (~ (e3 = op(e3, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1329, plain, (~ spl66_122 | spl66_42), inference(avatar_split_clause, [], [f481, f776, f1323])).
fof(f481, plain, ((e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f1328, plain, (~ spl66_122 | ~ spl66_42), inference(avatar_split_clause, [], [f482, f776, f1323])).
fof(f482, plain, (~ (e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1321, plain, (~ spl66_121 | spl66_42), inference(avatar_split_clause, [], [f477, f776, f1315])).
fof(f477, plain, ((e1 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f1320, plain, (~ spl66_121 | ~ spl66_42), inference(avatar_split_clause, [], [f478, f776, f1315])).
fof(f478, plain, (~ (e1 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1313, plain, (~ spl66_120 | spl66_42), inference(avatar_split_clause, [], [f473, f776, f1307])).
fof(f473, plain, ((e1 = op(e1, e1)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f1312, plain, (~ spl66_120 | ~ spl66_42), inference(avatar_split_clause, [], [f474, f776, f1307])).
fof(f474, plain, (~ (e1 = op(e1, e1)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1305, plain, (~ spl66_119 | spl66_42), inference(avatar_split_clause, [], [f469, f776, f1299])).
fof(f469, plain, ((e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f1304, plain, (~ spl66_119 | ~ spl66_42), inference(avatar_split_clause, [], [f470, f776, f1299])).
fof(f470, plain, (~ (e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f1297, plain, (~ spl66_118 | spl66_38), inference(avatar_split_clause, [], [f465, f759, f1291])).
fof(f465, plain, ((e1 = op(e1, e2)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f1296, plain, (~ spl66_118 | ~ spl66_43), inference(avatar_split_clause, [], [f466, f780, f1291])).
fof(f466, plain, (~ (e2 = op(e1, e1)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1295, plain, (~ spl66_118 | spl66_63), inference(avatar_split_clause, [], [f467, f865, f1291])).
fof(f467, plain, ((op(e0, e0) = e2) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1294, plain, (~ spl66_118 | ~ spl66_53), inference(avatar_split_clause, [], [f468, f823, f1291])).
fof(f468, plain, (~ (e0 = op(e0, e2)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1289, plain, (~ spl66_117 | spl66_38), inference(avatar_split_clause, [], [f461, f759, f1283])).
fof(f461, plain, ((e1 = op(e1, e2)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f1288, plain, (~ spl66_117 | ~ spl66_43), inference(avatar_split_clause, [], [f462, f780, f1283])).
fof(f462, plain, (~ (e2 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f1287, plain, (~ spl66_117 | spl66_43), inference(avatar_split_clause, [], [f463, f780, f1283])).
fof(f463, plain, ((e2 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f1279, plain, (~ spl66_116 | spl66_23), inference(avatar_split_clause, [], [f459, f695, f1275])).
fof(f459, plain, ((e2 = op(e2, e2)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f1278, plain, (~ spl66_116 | ~ spl66_23), inference(avatar_split_clause, [], [f460, f695, f1275])).
fof(f460, plain, (~ (e2 = op(e2, e2)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1273, plain, (~ spl66_115 | spl66_38), inference(avatar_split_clause, [], [f453, f759, f1267])).
fof(f453, plain, ((e1 = op(e1, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f1272, plain, (~ spl66_115 | ~ spl66_43), inference(avatar_split_clause, [], [f454, f780, f1267])).
fof(f454, plain, (~ (e2 = op(e1, e1)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1271, plain, (~ spl66_115 | spl66_3), inference(avatar_split_clause, [], [f455, f610, f1267])).
fof(f455, plain, ((e2 = op(e3, e3)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1270, plain, (~ spl66_115 | ~ spl66_8), inference(avatar_split_clause, [], [f456, f631, f1267])).
fof(f456, plain, (~ (e3 = op(e3, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1265, plain, (~ spl66_114 | spl66_34), inference(avatar_split_clause, [], [f449, f742, f1259])).
fof(f449, plain, ((e1 = op(e1, e3)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f1264, plain, (~ spl66_114 | ~ spl66_44), inference(avatar_split_clause, [], [f450, f784, f1259])).
fof(f450, plain, (~ (e3 = op(e1, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1263, plain, (~ spl66_114 | spl66_64), inference(avatar_split_clause, [], [f451, f869, f1259])).
fof(f451, plain, ((op(e0, e0) = e3) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1262, plain, (~ spl66_114 | ~ spl66_49), inference(avatar_split_clause, [], [f452, f806, f1259])).
fof(f452, plain, (~ (e0 = op(e0, e3)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1256, plain, (~ spl66_113 | ~ spl66_44), inference(avatar_split_clause, [], [f446, f784, f1251])).
fof(f446, plain, (~ (e3 = op(e1, e1)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f1255, plain, (~ spl66_113 | spl66_44), inference(avatar_split_clause, [], [f447, f784, f1251])).
fof(f447, plain, ((e3 = op(e1, e1)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f1249, plain, (~ spl66_112 | spl66_34), inference(avatar_split_clause, [], [f441, f742, f1243])).
fof(f441, plain, ((e1 = op(e1, e3)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP30), inference(nnf_transformation, [], [f84])).
fof(f1248, plain, (~ spl66_112 | ~ spl66_44), inference(avatar_split_clause, [], [f442, f784, f1243])).
fof(f442, plain, (~ (e3 = op(e1, e1)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1247, plain, (~ spl66_112 | spl66_24), inference(avatar_split_clause, [], [f443, f699, f1243])).
fof(f443, plain, ((e3 = op(e2, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1246, plain, (~ spl66_112 | ~ spl66_19), inference(avatar_split_clause, [], [f444, f678, f1243])).
fof(f444, plain, (~ (e2 = op(e2, e3)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1240, plain, (~ spl66_111 | ~ spl66_44), inference(avatar_split_clause, [], [f438, f784, f1235])).
fof(f438, plain, (~ (e3 = op(e1, e1)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(nnf_transformation, [], [f85])).
fof(f1239, plain, (~ spl66_111 | spl66_4), inference(avatar_split_clause, [], [f439, f614, f1235])).
fof(f439, plain, ((e3 = op(e3, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1238, plain, (~ spl66_111 | ~ spl66_4), inference(avatar_split_clause, [], [f440, f614, f1235])).
fof(f440, plain, (~ (e3 = op(e3, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1231, plain, (~ spl66_110 | spl66_61), inference(avatar_split_clause, [], [f435, f857, f1227])).
fof(f435, plain, ((e0 = op(e0, e0)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP32), inference(nnf_transformation, [], [f86])).
fof(f1230, plain, (~ spl66_110 | ~ spl66_61), inference(avatar_split_clause, [], [f436, f857, f1227])).
fof(f436, plain, (~ (e0 = op(e0, e0)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1225, plain, (~ spl66_109 | spl66_31), inference(avatar_split_clause, [], [f429, f729, f1219])).
fof(f429, plain, ((e2 = op(e2, e0)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP33), inference(nnf_transformation, [], [f87])).
fof(f1224, plain, (~ spl66_109 | ~ spl66_21), inference(avatar_split_clause, [], [f430, f687, f1219])).
fof(f430, plain, (~ (e0 = op(e2, e2)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1223, plain, (~ spl66_109 | spl66_41), inference(avatar_split_clause, [], [f431, f772, f1219])).
fof(f431, plain, ((e0 = op(e1, e1)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1222, plain, (~ spl66_109 | ~ spl66_46), inference(avatar_split_clause, [], [f432, f793, f1219])).
fof(f432, plain, (~ (e1 = op(e1, e0)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1217, plain, (~ spl66_108 | spl66_31), inference(avatar_split_clause, [], [f425, f729, f1211])).
fof(f425, plain, ((e2 = op(e2, e0)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP34), inference(nnf_transformation, [], [f88])).
fof(f1216, plain, (~ spl66_108 | ~ spl66_21), inference(avatar_split_clause, [], [f426, f687, f1211])).
fof(f426, plain, (~ (e0 = op(e2, e2)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f1215, plain, (~ spl66_108 | spl66_21), inference(avatar_split_clause, [], [f427, f687, f1211])).
fof(f427, plain, ((e0 = op(e2, e2)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f1209, plain, (~ spl66_107 | spl66_31), inference(avatar_split_clause, [], [f421, f729, f1203])).
fof(f421, plain, ((e2 = op(e2, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(nnf_transformation, [], [f89])).
fof(f1208, plain, (~ spl66_107 | ~ spl66_21), inference(avatar_split_clause, [], [f422, f687, f1203])).
fof(f422, plain, (~ (e0 = op(e2, e2)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1207, plain, (~ spl66_107 | spl66_1), inference(avatar_split_clause, [], [f423, f602, f1203])).
fof(f423, plain, ((e0 = op(e3, e3)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1206, plain, (~ spl66_107 | ~ spl66_16), inference(avatar_split_clause, [], [f424, f665, f1203])).
fof(f424, plain, (~ (e3 = op(e3, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1201, plain, (~ spl66_106 | spl66_27), inference(avatar_split_clause, [], [f417, f712, f1195])).
fof(f417, plain, ((e2 = op(e2, e1)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP36), inference(nnf_transformation, [], [f90])).
fof(f1200, plain, (~ spl66_106 | ~ spl66_22), inference(avatar_split_clause, [], [f418, f691, f1195])).
fof(f418, plain, (~ (e1 = op(e2, e2)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1199, plain, (~ spl66_106 | spl66_62), inference(avatar_split_clause, [], [f419, f861, f1195])).
fof(f419, plain, ((op(e0, e0) = e1) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1198, plain, (~ spl66_106 | ~ spl66_57), inference(avatar_split_clause, [], [f420, f840, f1195])).
fof(f420, plain, (~ (e0 = op(e0, e1)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1191, plain, (~ spl66_105 | spl66_42), inference(avatar_split_clause, [], [f415, f776, f1187])).
fof(f415, plain, ((e1 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP37), inference(nnf_transformation, [], [f91])).
fof(f1190, plain, (~ spl66_105 | ~ spl66_42), inference(avatar_split_clause, [], [f416, f776, f1187])).
fof(f416, plain, (~ (e1 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1185, plain, (~ spl66_104 | spl66_27), inference(avatar_split_clause, [], [f409, f712, f1179])).
fof(f409, plain, ((e2 = op(e2, e1)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP38), inference(nnf_transformation, [], [f92])).
fof(f1184, plain, (~ spl66_104 | ~ spl66_22), inference(avatar_split_clause, [], [f410, f691, f1179])).
fof(f410, plain, (~ (e1 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1183, plain, (~ spl66_104 | spl66_22), inference(avatar_split_clause, [], [f411, f691, f1179])).
fof(f411, plain, ((e1 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1182, plain, (~ spl66_104 | ~ spl66_27), inference(avatar_split_clause, [], [f412, f712, f1179])).
fof(f412, plain, (~ (e2 = op(e2, e1)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1177, plain, (~ spl66_103 | spl66_27), inference(avatar_split_clause, [], [f405, f712, f1171])).
fof(f405, plain, ((e2 = op(e2, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(nnf_transformation, [], [f93])).
fof(f1176, plain, (~ spl66_103 | ~ spl66_22), inference(avatar_split_clause, [], [f406, f691, f1171])).
fof(f406, plain, (~ (e1 = op(e2, e2)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1175, plain, (~ spl66_103 | spl66_2), inference(avatar_split_clause, [], [f407, f606, f1171])).
fof(f407, plain, ((e1 = op(e3, e3)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1174, plain, (~ spl66_103 | ~ spl66_12), inference(avatar_split_clause, [], [f408, f648, f1171])).
fof(f408, plain, (~ (e3 = op(e3, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1169, plain, (~ spl66_102 | spl66_23), inference(avatar_split_clause, [], [f401, f695, f1163])).
fof(f401, plain, ((e2 = op(e2, e2)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP40), inference(nnf_transformation, [], [f94])).
fof(f1168, plain, (~ spl66_102 | ~ spl66_23), inference(avatar_split_clause, [], [f402, f695, f1163])).
fof(f402, plain, (~ (e2 = op(e2, e2)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f1161, plain, (~ spl66_101 | spl66_23), inference(avatar_split_clause, [], [f397, f695, f1155])).
fof(f397, plain, ((e2 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP41), inference(nnf_transformation, [], [f95])).
fof(f1160, plain, (~ spl66_101 | ~ spl66_23), inference(avatar_split_clause, [], [f398, f695, f1155])).
fof(f398, plain, (~ (e2 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1153, plain, (~ spl66_100 | spl66_23), inference(avatar_split_clause, [], [f393, f695, f1147])).
fof(f393, plain, ((e2 = op(e2, e2)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP42), inference(nnf_transformation, [], [f96])).
fof(f1152, plain, (~ spl66_100 | ~ spl66_23), inference(avatar_split_clause, [], [f394, f695, f1147])).
fof(f394, plain, (~ (e2 = op(e2, e2)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1145, plain, (~ spl66_99 | spl66_23), inference(avatar_split_clause, [], [f389, f695, f1139])).
fof(f389, plain, ((e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(nnf_transformation, [], [f97])).
fof(f1144, plain, (~ spl66_99 | ~ spl66_23), inference(avatar_split_clause, [], [f390, f695, f1139])).
fof(f390, plain, (~ (e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1137, plain, (~ spl66_98 | spl66_19), inference(avatar_split_clause, [], [f385, f678, f1131])).
fof(f385, plain, ((e2 = op(e2, e3)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP44), inference(nnf_transformation, [], [f98])).
fof(f1136, plain, (~ spl66_98 | ~ spl66_24), inference(avatar_split_clause, [], [f386, f699, f1131])).
fof(f386, plain, (~ (e3 = op(e2, e2)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1135, plain, (~ spl66_98 | spl66_64), inference(avatar_split_clause, [], [f387, f869, f1131])).
fof(f387, plain, ((op(e0, e0) = e3) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1134, plain, (~ spl66_98 | ~ spl66_49), inference(avatar_split_clause, [], [f388, f806, f1131])).
fof(f388, plain, (~ (e0 = op(e0, e3)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1129, plain, (~ spl66_97 | spl66_19), inference(avatar_split_clause, [], [f381, f678, f1123])).
fof(f381, plain, ((e2 = op(e2, e3)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP45), inference(nnf_transformation, [], [f99])).
fof(f1128, plain, (~ spl66_97 | ~ spl66_24), inference(avatar_split_clause, [], [f382, f699, f1123])).
fof(f382, plain, (~ (e3 = op(e2, e2)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1127, plain, (~ spl66_97 | spl66_44), inference(avatar_split_clause, [], [f383, f784, f1123])).
fof(f383, plain, ((e3 = op(e1, e1)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1126, plain, (~ spl66_97 | ~ spl66_34), inference(avatar_split_clause, [], [f384, f742, f1123])).
fof(f384, plain, (~ (e1 = op(e1, e3)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1120, plain, (~ spl66_96 | ~ spl66_24), inference(avatar_split_clause, [], [f378, f699, f1115])).
fof(f378, plain, (~ (e3 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP46), inference(nnf_transformation, [], [f100])).
fof(f1119, plain, (~ spl66_96 | spl66_24), inference(avatar_split_clause, [], [f379, f699, f1115])).
fof(f379, plain, ((e3 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f1111, plain, (~ spl66_95 | spl66_4), inference(avatar_split_clause, [], [f375, f614, f1107])).
fof(f375, plain, ((e3 = op(e3, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(nnf_transformation, [], [f101])).
fof(f1110, plain, (~ spl66_95 | ~ spl66_4), inference(avatar_split_clause, [], [f376, f614, f1107])).
fof(f376, plain, (~ (e3 = op(e3, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1103, plain, (~ spl66_94 | spl66_61), inference(avatar_split_clause, [], [f371, f857, f1099])).
fof(f371, plain, ((e0 = op(e0, e0)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP48), inference(nnf_transformation, [], [f102])).
fof(f1102, plain, (~ spl66_94 | ~ spl66_61), inference(avatar_split_clause, [], [f372, f857, f1099])).
fof(f372, plain, (~ (e0 = op(e0, e0)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1097, plain, (~ spl66_93 | spl66_16), inference(avatar_split_clause, [], [f365, f665, f1091])).
fof(f365, plain, ((e3 = op(e3, e0)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP49), inference(nnf_transformation, [], [f103])).
fof(f1096, plain, (~ spl66_93 | ~ spl66_1), inference(avatar_split_clause, [], [f366, f602, f1091])).
fof(f366, plain, (~ (e0 = op(e3, e3)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1095, plain, (~ spl66_93 | spl66_41), inference(avatar_split_clause, [], [f367, f772, f1091])).
fof(f367, plain, ((e0 = op(e1, e1)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1094, plain, (~ spl66_93 | ~ spl66_46), inference(avatar_split_clause, [], [f368, f793, f1091])).
fof(f368, plain, (~ (e1 = op(e1, e0)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1089, plain, (~ spl66_92 | spl66_16), inference(avatar_split_clause, [], [f361, f665, f1083])).
fof(f361, plain, ((e3 = op(e3, e0)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP50), inference(nnf_transformation, [], [f104])).
fof(f1088, plain, (~ spl66_92 | ~ spl66_1), inference(avatar_split_clause, [], [f362, f602, f1083])).
fof(f362, plain, (~ (e0 = op(e3, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1087, plain, (~ spl66_92 | spl66_21), inference(avatar_split_clause, [], [f363, f687, f1083])).
fof(f363, plain, ((e0 = op(e2, e2)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1086, plain, (~ spl66_92 | ~ spl66_31), inference(avatar_split_clause, [], [f364, f729, f1083])).
fof(f364, plain, (~ (e2 = op(e2, e0)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1080, plain, (~ spl66_91 | ~ spl66_1), inference(avatar_split_clause, [], [f358, f602, f1075])).
fof(f358, plain, (~ (e0 = op(e3, e3)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(nnf_transformation, [], [f105])).
fof(f1079, plain, (~ spl66_91 | spl66_1), inference(avatar_split_clause, [], [f359, f602, f1075])).
fof(f359, plain, ((e0 = op(e3, e3)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f1073, plain, (~ spl66_90 | spl66_12), inference(avatar_split_clause, [], [f353, f648, f1067])).
fof(f353, plain, ((e3 = op(e3, e1)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP52), inference(nnf_transformation, [], [f106])).
fof(f1072, plain, (~ spl66_90 | ~ spl66_2), inference(avatar_split_clause, [], [f354, f606, f1067])).
fof(f354, plain, (~ (e1 = op(e3, e3)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1071, plain, (~ spl66_90 | spl66_62), inference(avatar_split_clause, [], [f355, f861, f1067])).
fof(f355, plain, ((op(e0, e0) = e1) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1070, plain, (~ spl66_90 | ~ spl66_57), inference(avatar_split_clause, [], [f356, f840, f1067])).
fof(f356, plain, (~ (e0 = op(e0, e1)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1063, plain, (~ spl66_89 | spl66_42), inference(avatar_split_clause, [], [f351, f776, f1059])).
fof(f351, plain, ((e1 = op(e1, e1)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP53), inference(nnf_transformation, [], [f107])).
fof(f1062, plain, (~ spl66_89 | ~ spl66_42), inference(avatar_split_clause, [], [f352, f776, f1059])).
fof(f352, plain, (~ (e1 = op(e1, e1)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f1057, plain, (~ spl66_88 | spl66_12), inference(avatar_split_clause, [], [f345, f648, f1051])).
fof(f345, plain, ((e3 = op(e3, e1)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP54), inference(nnf_transformation, [], [f108])).
fof(f1056, plain, (~ spl66_88 | ~ spl66_2), inference(avatar_split_clause, [], [f346, f606, f1051])).
fof(f346, plain, (~ (e1 = op(e3, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f1055, plain, (~ spl66_88 | spl66_22), inference(avatar_split_clause, [], [f347, f691, f1051])).
fof(f347, plain, ((e1 = op(e2, e2)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f1054, plain, (~ spl66_88 | ~ spl66_27), inference(avatar_split_clause, [], [f348, f712, f1051])).
fof(f348, plain, (~ (e2 = op(e2, e1)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f1049, plain, (~ spl66_87 | spl66_12), inference(avatar_split_clause, [], [f341, f648, f1043])).
fof(f341, plain, ((e3 = op(e3, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(nnf_transformation, [], [f109])).
fof(f1048, plain, (~ spl66_87 | ~ spl66_2), inference(avatar_split_clause, [], [f342, f606, f1043])).
fof(f342, plain, (~ (e1 = op(e3, e3)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f1047, plain, (~ spl66_87 | spl66_2), inference(avatar_split_clause, [], [f343, f606, f1043])).
fof(f343, plain, ((e1 = op(e3, e3)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f1041, plain, (~ spl66_86 | spl66_8), inference(avatar_split_clause, [], [f337, f631, f1035])).
fof(f337, plain, ((e3 = op(e3, e2)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP56), inference(nnf_transformation, [], [f110])).
fof(f1040, plain, (~ spl66_86 | ~ spl66_3), inference(avatar_split_clause, [], [f338, f610, f1035])).
fof(f338, plain, (~ (e2 = op(e3, e3)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f1039, plain, (~ spl66_86 | spl66_63), inference(avatar_split_clause, [], [f339, f865, f1035])).
fof(f339, plain, ((op(e0, e0) = e2) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f1038, plain, (~ spl66_86 | ~ spl66_53), inference(avatar_split_clause, [], [f340, f823, f1035])).
fof(f340, plain, (~ (e0 = op(e0, e2)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f1033, plain, (~ spl66_85 | spl66_8), inference(avatar_split_clause, [], [f333, f631, f1027])).
fof(f333, plain, ((e3 = op(e3, e2)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP57), inference(nnf_transformation, [], [f111])).
fof(f1032, plain, (~ spl66_85 | ~ spl66_3), inference(avatar_split_clause, [], [f334, f610, f1027])).
fof(f334, plain, (~ (e2 = op(e3, e3)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f1031, plain, (~ spl66_85 | spl66_43), inference(avatar_split_clause, [], [f335, f780, f1027])).
fof(f335, plain, ((e2 = op(e1, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f1030, plain, (~ spl66_85 | ~ spl66_38), inference(avatar_split_clause, [], [f336, f759, f1027])).
fof(f336, plain, (~ (e1 = op(e1, e2)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f1025, plain, (~ spl66_84 | spl66_8), inference(avatar_split_clause, [], [f329, f631, f1019])).
fof(f329, plain, ((e3 = op(e3, e2)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP58), inference(nnf_transformation, [], [f112])).
fof(f1023, plain, (~ spl66_84 | spl66_23), inference(avatar_split_clause, [], [f331, f695, f1019])).
fof(f331, plain, ((e2 = op(e2, e2)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f1022, plain, (~ spl66_84 | ~ spl66_23), inference(avatar_split_clause, [], [f332, f695, f1019])).
fof(f332, plain, (~ (e2 = op(e2, e2)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f1016, plain, (~ spl66_83 | ~ spl66_3), inference(avatar_split_clause, [], [f326, f610, f1011])).
fof(f326, plain, (~ (e2 = op(e3, e3)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(nnf_transformation, [], [f113])).
fof(f1015, plain, (~ spl66_83 | spl66_3), inference(avatar_split_clause, [], [f327, f610, f1011])).
fof(f327, plain, ((e2 = op(e3, e3)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f1009, plain, (~ spl66_82 | spl66_4), inference(avatar_split_clause, [], [f321, f614, f1003])).
fof(f321, plain, ((e3 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP60), inference(nnf_transformation, [], [f114])).
fof(f1008, plain, (~ spl66_82 | ~ spl66_4), inference(avatar_split_clause, [], [f322, f614, f1003])).
fof(f322, plain, (~ (e3 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f1001, plain, (~ spl66_81 | spl66_4), inference(avatar_split_clause, [], [f317, f614, f995])).
fof(f317, plain, ((e3 = op(e3, e3)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP61), inference(nnf_transformation, [], [f115])).
fof(f1000, plain, (~ spl66_81 | ~ spl66_4), inference(avatar_split_clause, [], [f318, f614, f995])).
fof(f318, plain, (~ (e3 = op(e3, e3)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f993, plain, (~ spl66_80 | spl66_4), inference(avatar_split_clause, [], [f313, f614, f987])).
fof(f313, plain, ((e3 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP62), inference(nnf_transformation, [], [f116])).
fof(f992, plain, (~ spl66_80 | ~ spl66_4), inference(avatar_split_clause, [], [f314, f614, f987])).
fof(f314, plain, (~ (e3 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f985, plain, (~ spl66_75 | spl66_79), inference(avatar_split_clause, [], [f305, f982, f960])).
fof(f305, plain, ((e0 = op(e0, op(e0, e0))) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0)))) | ~ sP63), inference(nnf_transformation, [], [f117])).
fof(f980, plain, (~ spl66_75 | spl66_78), inference(avatar_split_clause, [], [f306, f977, f960])).
fof(f306, plain, ((e1 = op(e0, op(e0, e1))) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f975, plain, (~ spl66_75 | spl66_77), inference(avatar_split_clause, [], [f307, f972, f960])).
fof(f307, plain, ((e2 = op(e0, op(e0, e2))) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f970, plain, (~ spl66_75 | spl66_76), inference(avatar_split_clause, [], [f308, f967, f960])).
fof(f308, plain, ((e3 = op(e0, op(e0, e3))) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f965, plain, (~ spl66_75 | ~ spl66_46 | spl66_41), inference(avatar_split_clause, [], [f310, f772, f793, f960])).
fof(f310, plain, ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f964, plain, (~ spl66_75 | ~ spl66_31 | spl66_21), inference(avatar_split_clause, [], [f311, f687, f729, f960])).
fof(f311, plain, ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f963, plain, (~ spl66_75 | ~ spl66_16 | spl66_1), inference(avatar_split_clause, [], [f312, f602, f665, f960])).
fof(f312, plain, ((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f958, plain, (~ spl66_70 | spl66_74), inference(avatar_split_clause, [], [f297, f955, f933])).
fof(f297, plain, ((e0 = op(e1, op(e1, e0))) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0)))) | ~ sP64), inference(nnf_transformation, [], [f118])).
fof(f953, plain, (~ spl66_70 | spl66_73), inference(avatar_split_clause, [], [f298, f950, f933])).
fof(f298, plain, ((e1 = op(e1, op(e1, e1))) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f948, plain, (~ spl66_70 | spl66_72), inference(avatar_split_clause, [], [f299, f945, f933])).
fof(f299, plain, ((e2 = op(e1, op(e1, e2))) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f943, plain, (~ spl66_70 | spl66_71), inference(avatar_split_clause, [], [f300, f940, f933])).
fof(f300, plain, ((e3 = op(e1, op(e1, e3))) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f938, plain, (~ spl66_70 | ~ spl66_57 | spl66_62), inference(avatar_split_clause, [], [f301, f861, f840, f933])).
fof(f301, plain, ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f937, plain, (~ spl66_70 | ~ spl66_27 | spl66_22), inference(avatar_split_clause, [], [f303, f691, f712, f933])).
fof(f303, plain, ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f936, plain, (~ spl66_70 | ~ spl66_12 | spl66_2), inference(avatar_split_clause, [], [f304, f606, f648, f933])).
fof(f304, plain, ((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f931, plain, (~ spl66_65 | spl66_69), inference(avatar_split_clause, [], [f289, f928, f906])).
fof(f289, plain, ((e0 = op(e2, op(e2, e0))) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0)))) | ~ sP65), inference(nnf_transformation, [], [f119])).
fof(f926, plain, (~ spl66_65 | spl66_68), inference(avatar_split_clause, [], [f290, f923, f906])).
fof(f290, plain, ((e1 = op(e2, op(e2, e1))) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f921, plain, (~ spl66_65 | spl66_67), inference(avatar_split_clause, [], [f291, f918, f906])).
fof(f291, plain, ((e2 = op(e2, op(e2, e2))) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f916, plain, (~ spl66_65 | spl66_66), inference(avatar_split_clause, [], [f292, f913, f906])).
fof(f292, plain, ((e3 = op(e2, op(e2, e3))) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f911, plain, (~ spl66_65 | ~ spl66_53 | spl66_63), inference(avatar_split_clause, [], [f293, f865, f823, f906])).
fof(f293, plain, ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f910, plain, (~ spl66_65 | ~ spl66_38 | spl66_43), inference(avatar_split_clause, [], [f294, f780, f759, f906])).
fof(f294, plain, ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f909, plain, (~ spl66_65 | ~ spl66_8 | spl66_3), inference(avatar_split_clause, [], [f296, f610, f631, f906])).
fof(f296, plain, ((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f904, plain, (spl66_61 | spl66_57 | spl66_53 | spl66_49), inference(avatar_split_clause, [], [f203, f806, f823, f840, f857])).
fof(f203, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax2)).
fof(f902, plain, (spl66_62 | spl66_58 | spl66_54 | spl66_50), inference(avatar_split_clause, [], [f205, f810, f827, f844, f861])).
fof(f205, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f901, plain, (spl66_62 | spl66_46 | spl66_30 | spl66_14), inference(avatar_split_clause, [], [f206, f657, f725, f793, f861])).
fof(f206, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f900, plain, (spl66_63 | spl66_59 | spl66_55 | spl66_51), inference(avatar_split_clause, [], [f207, f814, f831, f848, f865])).
fof(f207, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f899, plain, (spl66_63 | spl66_47 | spl66_31 | spl66_15), inference(avatar_split_clause, [], [f208, f661, f729, f797, f865])).
fof(f208, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f898, plain, (spl66_64 | spl66_60 | spl66_56 | spl66_52), inference(avatar_split_clause, [], [f209, f818, f835, f852, f869])).
fof(f209, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f897, plain, (spl66_64 | spl66_48 | spl66_32 | spl66_16), inference(avatar_split_clause, [], [f210, f665, f733, f801, f869])).
fof(f210, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f896, plain, (spl66_45 | spl66_41 | spl66_37 | spl66_33), inference(avatar_split_clause, [], [f211, f738, f755, f772, f789])).
fof(f211, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f895, plain, (spl66_57 | spl66_41 | spl66_25 | spl66_9), inference(avatar_split_clause, [], [f212, f636, f704, f772, f840])).
fof(f212, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f894, plain, (spl66_46 | spl66_42 | spl66_38 | spl66_34), inference(avatar_split_clause, [], [f213, f742, f759, f776, f793])).
fof(f213, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f891, plain, (spl66_59 | spl66_43 | spl66_27 | spl66_11), inference(avatar_split_clause, [], [f216, f644, f712, f780, f848])).
fof(f216, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f890, plain, (spl66_48 | spl66_44 | spl66_40 | spl66_36), inference(avatar_split_clause, [], [f217, f750, f767, f784, f801])).
fof(f217, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f889, plain, (spl66_60 | spl66_44 | spl66_28 | spl66_12), inference(avatar_split_clause, [], [f218, f648, f716, f784, f852])).
fof(f218, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f887, plain, (spl66_53 | spl66_37 | spl66_21 | spl66_5), inference(avatar_split_clause, [], [f220, f619, f687, f755, f823])).
fof(f220, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f886, plain, (spl66_30 | spl66_26 | spl66_22 | spl66_18), inference(avatar_split_clause, [], [f221, f674, f691, f708, f725])).
fof(f221, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f885, plain, (spl66_54 | spl66_38 | spl66_22 | spl66_6), inference(avatar_split_clause, [], [f222, f623, f691, f759, f827])).
fof(f222, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f884, plain, (spl66_31 | spl66_27 | spl66_23 | spl66_19), inference(avatar_split_clause, [], [f223, f678, f695, f712, f729])).
fof(f223, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f883, plain, (spl66_55 | spl66_39 | spl66_23 | spl66_7), inference(avatar_split_clause, [], [f224, f627, f695, f763, f831])).
fof(f224, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f882, plain, (spl66_32 | spl66_28 | spl66_24 | spl66_20), inference(avatar_split_clause, [], [f225, f682, f699, f716, f733])).
fof(f225, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f881, plain, (spl66_56 | spl66_40 | spl66_24 | spl66_8), inference(avatar_split_clause, [], [f226, f631, f699, f767, f835])).
fof(f226, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f880, plain, (spl66_13 | spl66_9 | spl66_5 | spl66_1), inference(avatar_split_clause, [], [f227, f602, f619, f636, f653])).
fof(f227, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f879, plain, (spl66_49 | spl66_33 | spl66_17 | spl66_1), inference(avatar_split_clause, [], [f228, f602, f670, f738, f806])).
fof(f228, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f878, plain, (spl66_14 | spl66_10 | spl66_6 | spl66_2), inference(avatar_split_clause, [], [f229, f606, f623, f640, f657])).
fof(f229, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f877, plain, (spl66_50 | spl66_34 | spl66_18 | spl66_2), inference(avatar_split_clause, [], [f230, f606, f674, f742, f810])).
fof(f230, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f876, plain, (spl66_15 | spl66_11 | spl66_7 | spl66_3), inference(avatar_split_clause, [], [f231, f610, f627, f644, f661])).
fof(f231, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f875, plain, (spl66_51 | spl66_35 | spl66_19 | spl66_3), inference(avatar_split_clause, [], [f232, f610, f678, f746, f814])).
fof(f232, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f874, plain, (spl66_16 | spl66_12 | spl66_8 | spl66_4), inference(avatar_split_clause, [], [f233, f614, f631, f648, f665])).
fof(f233, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f873, plain, (spl66_52 | spl66_36 | spl66_20 | spl66_4), inference(avatar_split_clause, [], [f234, f614, f682, f750, f818])).
fof(f234, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f872, plain, (spl66_61 | spl66_62 | spl66_63 | spl66_64), inference(avatar_split_clause, [], [f187, f869, f865, f861, f857])).
fof(f187, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG155+1.p', ax1)).
fof(f855, plain, (spl66_57 | spl66_58 | spl66_59 | spl66_60), inference(avatar_split_clause, [], [f188, f852, f848, f844, f840])).
fof(f188, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f838, plain, (spl66_53 | spl66_54 | spl66_55 | spl66_56), inference(avatar_split_clause, [], [f189, f835, f831, f827, f823])).
fof(f189, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f821, plain, (spl66_49 | spl66_50 | spl66_51 | spl66_52), inference(avatar_split_clause, [], [f190, f818, f814, f810, f806])).
fof(f190, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f804, plain, (spl66_45 | spl66_46 | spl66_47 | spl66_48), inference(avatar_split_clause, [], [f191, f801, f797, f793, f789])).
fof(f191, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f787, plain, (spl66_41 | spl66_42 | spl66_43 | spl66_44), inference(avatar_split_clause, [], [f192, f784, f780, f776, f772])).
fof(f192, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f770, plain, (spl66_37 | spl66_38 | spl66_39 | spl66_40), inference(avatar_split_clause, [], [f193, f767, f763, f759, f755])).
fof(f193, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f753, plain, (spl66_33 | spl66_34 | spl66_35 | spl66_36), inference(avatar_split_clause, [], [f194, f750, f746, f742, f738])).
fof(f194, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f736, plain, (spl66_29 | spl66_30 | spl66_31 | spl66_32), inference(avatar_split_clause, [], [f195, f733, f729, f725, f721])).
fof(f195, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f719, plain, (spl66_25 | spl66_26 | spl66_27 | spl66_28), inference(avatar_split_clause, [], [f196, f716, f712, f708, f704])).
fof(f196, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f702, plain, (spl66_21 | spl66_22 | spl66_23 | spl66_24), inference(avatar_split_clause, [], [f197, f699, f695, f691, f687])).
fof(f197, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f685, plain, (spl66_17 | spl66_18 | spl66_19 | spl66_20), inference(avatar_split_clause, [], [f198, f682, f678, f674, f670])).
fof(f198, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f668, plain, (spl66_13 | spl66_14 | spl66_15 | spl66_16), inference(avatar_split_clause, [], [f199, f665, f661, f657, f653])).
fof(f199, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f651, plain, (spl66_9 | spl66_10 | spl66_11 | spl66_12), inference(avatar_split_clause, [], [f200, f648, f644, f640, f636])).
fof(f200, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f634, plain, (spl66_5 | spl66_6 | spl66_7 | spl66_8), inference(avatar_split_clause, [], [f201, f631, f627, f623, f619])).
fof(f201, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f617, plain, (spl66_1 | spl66_2 | spl66_3 | spl66_4), inference(avatar_split_clause, [], [f202, f614, f610, f606, f602])).
fof(f202, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).