fof(f5733, plain, $false, inference(avatar_sat_refutation, [], [f532, f549, f566, f583, f600, f617, f634, f651, f668, f685, f702, f719, f736, f753, f770, f787, f788, f789, f790, f791, f792, f793, f794, f795, f796, f797, f798, f799, f800, f801, f802, f803, f804, f805, f806, f807, f808, f809, f810, f811, f812, f813, f814, f815, f816, f817, f818, f819, f824, f826, f827, f828, f833, f834, f837, f842, f843, f846, f852, f853, f854, f855, f861, f864, f869, f870, f872, f873, f878, f879, f880, f882, f888, f889, f890, f891, f896, f898, f899, f900, f907, f909, f914, f915, f916, f918, f924, f925, f926, f927, f932, f934, f935, f936, f941, f942, f944, f945, f953, f954, f963, f968, f977, f982, f991, f996, f1005, f1006, f1016, f1025, f1030, f1039, f1044, f1053, f1054, f1063, f1064, f1074, f1083, f1088, f1097, f1098, f1107, f1108, f1117, f1118, f1128, f1133, f1134, f1135, f1136, f1137, f1138, f1143, f1144, f1145, f1146, f1147, f1148, f1153, f1154, f1155, f1156, f1157, f1158, f1163, f1164, f1165, f1166, f1167, f1168, f1174, f1178, f1183, f1184, f1185, f1186, f1187, f1188, f1193, f1194, f1195, f1196, f1197, f1198, f1203, f1204, f1205, f1206, f1207, f1208, f1213, f1214, f1215, f1216, f1217, f1218, f1225, f1228, f1233, f1234, f1236, f1237, f1238, f1243, f1244, f1245, f1246, f1247, f1248, f1253, f1254, f1255, f1256, f1257, f1258, f1263, f1264, f1265, f1266, f1267, f1268, f1276, f1278, f1279, f1283, f1289, f1290, f1291, f1292, f1293, f1295, f1318, f1324, f1325, f1343, f1348, f1353, f1354, f1355, f1364, f1373, f1378, f1383, f1384, f1385, f1403, f1408, f1413, f1414, f1415, f1431, f1432, f1433, f1461, f1470, f1481, f1495, f1513, f1534, f1535, f1539, f1542, f1546, f1549, f1552, f1558, f1561, f1564, f1571, f1584, f1587, f1603, f1618, f1624, f1628, f1631, f1635, f1663, f1666, f1713, f1716, f1735, f1736, f1738, f1757, f1760, f1767, f1772, f1787, f1794, f1822, f1825, f1838, f1852, f1883, f1915, f1919, f1934, f1942, f1957, f1968, f1970, f2019, f2021, f2031, f2047, f2060, f2114, f2121, f2137, f2154, f2155, f2164, f2171, f2178, f2180, f2204, f2217, f2225, f2234, f2238, f2260, f2270, f2274, f2286, f2292, f2311, f2319, f2324, f2336, f2337, f2347, f2352, f2358, f2364, f2386, f2393, f2416, f2441, f2458, f2466, f2473, f2476, f2483, f2486, f2523, f2549, f2552, f2556, f2558, f2567, f2581, f2595, f2608, f2627, f2644, f2661, f2693, f2697, f2723, f2736, f2752, f2756, f2813, f2846, f2847, f2852, f2858, f2861, f2870, f2873, f2884, f2901, f2915, f2935, f2962, f2965, f2972, f2976, f2980, f2990, f3026, f3031, f3039, f3041, f3062, f3064, f3069, f3070, f3072, f3087, f3096, f3108, f3190, f3196, f3206, f3209, f3221, f3266, f3274, f3276, f3283, f3297, f3300, f3308, f3336, f3362, f3392, f3417, f3418, f3424, f3435, f3453, f3501, f3509, f3512, f3528, f3534, f3545, f3552, f3553, f3554, f3568, f3591, f3604, f3629, f3637, f3647, f3654, f3662, f3679, f3684, f3719, f3720, f3751, f3769, f3784, f3859, f3860, f3869, f3871, f3884, f3888, f3890, f3899, f3906, f3919, f3927, f3928, f3958, f3981, f3993, f4005, f4043, f4050, f4051, f4087, f4090, f4097, f4103, f4104, f4107, f4117, f4121, f4123, f4134, f4136, f4147, f4148, f4159, f4176, f4185, f4190, f4192, f4202, f4205, f4207, f4209, f4212, f4225, f4232, f4246, f4247, f4252, f4261, f4273, f4280, f4308, f4313, f4315, f4328, f4349, f4356, f4365, f4376, f4377, f4378, f4385, f4411, f4442, f4443, f4452, f4453, f4462, f4463, f4464, f4465, f4470, f4471, f4483, f4484, f4486, f4493, f4498, f4502, f4514, f4527, f4529, f4531, f4559, f4569, f4570, f4573, f4579, f4588, f4597, f4601, f4611, f4618, f4626, f4641, f4661, f4662, f4673, f4682, f4683, f4698, f4715, f4730, f4747, f4749, f4762, f4772, f4773, f4780, f4781, f4789, f4800, f4801, f4802, f4807, f4809, f4817, f4821, f4829, f4853, f4873, f4892, f4893, f4905, f4919, f4927, f4935, f4939, f4942, f4944, f4953, f4954, f4955, f4963, f4972, f4973, f4981, f4987, f5004, f5009, f5020, f5026, f5051, f5053, f5055, f5058, f5060, f5062, f5066, f5068, f5082, f5084, f5094, f5095, f5096, f5109, f5113, f5118, f5120, f5131, f5133, f5139, f5143, f5159, f5161, f5169, f5172, f5174, f5179, f5185, f5186, f5195, f5205, f5206, f5221, f5223, f5232, f5249, f5254, f5257, f5283, f5284, f5297, f5307, f5308, f5312, f5318, f5337, f5352, f5354, f5361, f5364, f5379, f5380, f5387, f5388, f5398, f5400, f5410, f5411, f5423, f5424, f5425, f5434, f5439, f5445, f5462, f5477, f5483, f5485, f5500, f5502, f5514, f5515, f5516, f5527, f5528, f5529, f5539, f5540, f5550, f5560, f5583, f5589, f5590, f5597, f5598, f5601, f5613, f5634, f5636, f5638, f5640, f5642, f5655, f5656, f5669, f5670, f5678, f5684, f5686, f5695, f5696, f5703, f5708, f5712, f5724, f5727])).
fof(f5727, plain, (~ spl45_4 | ~ spl45_8), inference(avatar_split_clause, [], [f5725, f546, f529])).
fof(f529, plain, (spl45_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_4])])).
fof(f546, plain, (spl45_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_8])])).
fof(f5725, plain, (~ (e3 = op(e3, e3)) | ~ spl45_8), inference(backward_demodulation, [], [f241, f548])).
fof(f548, plain, ((e3 = op(e3, e2)) | ~ spl45_8), inference(avatar_component_clause, [], [f546])).
fof(f241, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax3)).
fof(f5724, plain, (~ spl45_7 | ~ spl45_11), inference(avatar_split_clause, [], [f5720, f559, f542])).
fof(f542, plain, (spl45_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_7])])).
fof(f559, plain, (spl45_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_11])])).
fof(f5720, plain, (~ (e2 = op(e3, e2)) | ~ spl45_11), inference(backward_demodulation, [], [f239, f561])).
fof(f561, plain, ((e2 = op(e3, e1)) | ~ spl45_11), inference(avatar_component_clause, [], [f559])).
fof(f239, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5712, plain, (~ spl45_24 | ~ spl45_32), inference(avatar_split_clause, [], [f5710, f648, f614])).
fof(f614, plain, (spl45_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_24])])).
fof(f648, plain, (spl45_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_32])])).
fof(f5710, plain, (~ (e3 = op(e2, e2)) | ~ spl45_32), inference(backward_demodulation, [], [f231, f650])).
fof(f650, plain, ((e3 = op(e2, e0)) | ~ spl45_32), inference(avatar_component_clause, [], [f648])).
fof(f231, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f5708, plain, (~ spl45_4 | ~ spl45_36), inference(avatar_split_clause, [], [f5705, f665, f529])).
fof(f665, plain, (spl45_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_36])])).
fof(f5705, plain, (~ (e3 = op(e3, e3)) | ~ spl45_36), inference(backward_demodulation, [], [f216, f667])).
fof(f667, plain, ((e3 = op(e1, e3)) | ~ spl45_36), inference(avatar_component_clause, [], [f665])).
fof(f216, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5703, plain, (~ spl45_34 | ~ spl45_38), inference(avatar_split_clause, [], [f5699, f674, f657])).
fof(f657, plain, (spl45_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_34])])).
fof(f674, plain, (spl45_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_38])])).
fof(f5699, plain, (~ (e1 = op(e1, e3)) | ~ spl45_38), inference(backward_demodulation, [], [f229, f676])).
fof(f676, plain, ((e1 = op(e1, e2)) | ~ spl45_38), inference(avatar_component_clause, [], [f674])).
fof(f229, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f5696, plain, (~ spl45_37 | ~ spl45_41), inference(avatar_split_clause, [], [f5688, f687, f670])).
fof(f670, plain, (spl45_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_37])])).
fof(f687, plain, (spl45_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_41])])).
fof(f5688, plain, (~ (e0 = op(e1, e2)) | ~ spl45_41), inference(backward_demodulation, [], [f227, f689])).
fof(f689, plain, ((e0 = op(e1, e1)) | ~ spl45_41), inference(avatar_component_clause, [], [f687])).
fof(f227, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f5695, plain, (~ spl45_9 | ~ spl45_41), inference(avatar_split_clause, [], [f5687, f687, f551])).
fof(f551, plain, (spl45_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_9])])).
fof(f5687, plain, (~ (e0 = op(e3, e1)) | ~ spl45_41), inference(backward_demodulation, [], [f204, f689])).
fof(f204, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f5686, plain, (spl45_12 | ~ spl45_50 | ~ spl45_109), inference(avatar_contradiction_clause, [], [f5685])).
fof(f5685, plain, ($false | (spl45_12 | ~ spl45_50 | ~ spl45_109)), inference(subsumption_resolution, [], [f5683, f564])).
fof(f564, plain, (~ (e3 = op(e3, e1)) | spl45_12), inference(avatar_component_clause, [], [f563])).
fof(f563, plain, (spl45_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_12])])).
fof(f5683, plain, ((e3 = op(e3, e1)) | (~ spl45_50 | ~ spl45_109)), inference(backward_demodulation, [], [f1096, f727])).
fof(f727, plain, ((e1 = op(e0, e3)) | ~ spl45_50), inference(avatar_component_clause, [], [f725])).
fof(f725, plain, (spl45_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_50])])).
fof(f1096, plain, ((e3 = op(e3, op(e0, e3))) | ~ spl45_109), inference(avatar_component_clause, [], [f1094])).
fof(f1094, plain, (spl45_109 <=> (e3 = op(e3, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl45_109])])).
fof(f5684, plain, (~ spl45_34 | ~ spl45_50), inference(avatar_split_clause, [], [f5679, f725, f657])).
fof(f5679, plain, (~ (e1 = op(e1, e3)) | ~ spl45_50), inference(backward_demodulation, [], [f212, f727])).
fof(f212, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f5678, plain, (~ spl45_7 | ~ spl45_55), inference(avatar_split_clause, [], [f5673, f746, f542])).
fof(f746, plain, (spl45_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_55])])).
fof(f5673, plain, (~ (e2 = op(e3, e2)) | ~ spl45_55), inference(backward_demodulation, [], [f208, f748])).
fof(f748, plain, ((e2 = op(e0, e2)) | ~ spl45_55), inference(avatar_component_clause, [], [f746])).
fof(f208, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5670, plain, (~ spl45_34 | ~ spl45_60 | spl45_113), inference(avatar_split_clause, [], [f5668, f1114, f767, f657])).
fof(f767, plain, (spl45_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_60])])).
fof(f1114, plain, (spl45_113 <=> (e1 = op(e1, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl45_113])])).
fof(f5668, plain, (~ (e1 = op(e1, e3)) | (~ spl45_60 | spl45_113)), inference(backward_demodulation, [], [f1115, f769])).
fof(f769, plain, ((e3 = op(e0, e1)) | ~ spl45_60), inference(avatar_component_clause, [], [f767])).
fof(f1115, plain, (~ (e1 = op(e1, op(e0, e1))) | spl45_113), inference(avatar_component_clause, [], [f1114])).
fof(f5669, plain, (~ spl45_44 | ~ spl45_60), inference(avatar_split_clause, [], [f5664, f767, f699])).
fof(f699, plain, (spl45_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_44])])).
fof(f5664, plain, (~ (e3 = op(e1, e1)) | ~ spl45_60), inference(backward_demodulation, [], [f200, f769])).
fof(f200, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f5656, plain, (~ spl45_49 | ~ spl45_61), inference(avatar_split_clause, [], [f5646, f772, f721])).
fof(f721, plain, (spl45_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_49])])).
fof(f772, plain, (spl45_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_61])])).
fof(f5646, plain, (~ (e0 = op(e0, e3)) | ~ spl45_61), inference(backward_demodulation, [], [f220, f774])).
fof(f774, plain, ((e0 = op(e0, e0)) | ~ spl45_61), inference(avatar_component_clause, [], [f772])).
fof(f220, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f5655, plain, (~ spl45_29 | ~ spl45_61), inference(avatar_split_clause, [], [f5643, f772, f636])).
fof(f636, plain, (spl45_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_29])])).
fof(f5643, plain, (~ (e0 = op(e2, e0)) | ~ spl45_61), inference(backward_demodulation, [], [f195, f774])).
fof(f195, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f5642, plain, (~ spl45_7 | ~ spl45_19 | spl45_82), inference(avatar_split_clause, [], [f5641, f965, f593, f542])).
fof(f593, plain, (spl45_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_19])])).
fof(f965, plain, (spl45_82 <=> (op(e2, e3) = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_82])])).
fof(f5641, plain, (~ (e2 = op(e3, e2)) | (~ spl45_19 | spl45_82)), inference(forward_demodulation, [], [f966, f595])).
fof(f595, plain, ((e2 = op(e2, e3)) | ~ spl45_19), inference(avatar_component_clause, [], [f593])).
fof(f966, plain, (~ (op(e2, e3) = op(e3, e2)) | spl45_82), inference(avatar_component_clause, [], [f965])).
fof(f5640, plain, (spl45_50 | ~ spl45_14 | ~ spl45_88), inference(avatar_split_clause, [], [f5639, f993, f572, f725])).
fof(f572, plain, (spl45_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_14])])).
fof(f993, plain, (spl45_88 <=> (op(e0, e3) = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_88])])).
fof(f5639, plain, ((e1 = op(e0, e3)) | (~ spl45_14 | ~ spl45_88)), inference(forward_demodulation, [], [f995, f574])).
fof(f574, plain, ((e1 = op(e3, e0)) | ~ spl45_14), inference(avatar_component_clause, [], [f572])).
fof(f995, plain, ((op(e0, e3) = op(e3, e0)) | ~ spl45_88), inference(avatar_component_clause, [], [f993])).
fof(f5638, plain, (spl45_8 | ~ spl45_19 | ~ spl45_90), inference(avatar_split_clause, [], [f5637, f1002, f593, f546])).
fof(f1002, plain, (spl45_90 <=> (e3 = op(e3, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl45_90])])).
fof(f5637, plain, ((e3 = op(e3, e2)) | (~ spl45_19 | ~ spl45_90)), inference(forward_demodulation, [], [f1004, f595])).
fof(f1004, plain, ((e3 = op(e3, op(e2, e3))) | ~ spl45_90), inference(avatar_component_clause, [], [f1002])).
fof(f5636, plain, (spl45_38 | ~ spl45_26 | ~ spl45_95), inference(avatar_split_clause, [], [f5635, f1027, f623, f674])).
fof(f623, plain, (spl45_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_26])])).
fof(f1027, plain, (spl45_95 <=> (op(e1, e2) = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_95])])).
fof(f5635, plain, ((e1 = op(e1, e2)) | (~ spl45_26 | ~ spl45_95)), inference(forward_demodulation, [], [f1029, f625])).
fof(f625, plain, ((e1 = op(e2, e1)) | ~ spl45_26), inference(avatar_component_clause, [], [f623])).
fof(f1029, plain, ((op(e1, e2) = op(e2, e1)) | ~ spl45_95), inference(avatar_component_clause, [], [f1027])).
fof(f5634, plain, (~ spl45_59 | ~ spl45_47 | spl45_107), inference(avatar_split_clause, [], [f5633, f1085, f712, f763])).
fof(f763, plain, (spl45_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_59])])).
fof(f712, plain, (spl45_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_47])])).
fof(f1085, plain, (spl45_107 <=> (op(e0, e1) = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_107])])).
fof(f5633, plain, (~ (e2 = op(e0, e1)) | (~ spl45_47 | spl45_107)), inference(forward_demodulation, [], [f1086, f714])).
fof(f714, plain, ((e2 = op(e1, e0)) | ~ spl45_47), inference(avatar_component_clause, [], [f712])).
fof(f1086, plain, (~ (op(e0, e1) = op(e1, e0)) | spl45_107), inference(avatar_component_clause, [], [f1085])).
fof(f5613, plain, (~ spl45_10 | ~ spl45_14), inference(avatar_split_clause, [], [f5612, f572, f555])).
fof(f555, plain, (spl45_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_10])])).
fof(f5612, plain, (~ (e1 = op(e3, e1)) | ~ spl45_14), inference(forward_demodulation, [], [f236, f574])).
fof(f236, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f5601, plain, (~ spl45_19 | spl45_23 | ~ spl45_81 | ~ spl45_82), inference(avatar_contradiction_clause, [], [f5600])).
fof(f5600, plain, ($false | (~ spl45_19 | spl45_23 | ~ spl45_81 | ~ spl45_82)), inference(subsumption_resolution, [], [f5595, f611])).
fof(f611, plain, (~ (e2 = op(e2, e2)) | spl45_23), inference(avatar_component_clause, [], [f610])).
fof(f610, plain, (spl45_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_23])])).
fof(f5595, plain, ((e2 = op(e2, e2)) | (~ spl45_19 | ~ spl45_81 | ~ spl45_82)), inference(backward_demodulation, [], [f5504, f595])).
fof(f5504, plain, ((e2 = op(e2, op(e2, e3))) | (~ spl45_81 | ~ spl45_82)), inference(forward_demodulation, [], [f962, f967])).
fof(f967, plain, ((op(e2, e3) = op(e3, e2)) | ~ spl45_82), inference(avatar_component_clause, [], [f965])).
fof(f962, plain, ((e2 = op(e2, op(e3, e2))) | ~ spl45_81), inference(avatar_component_clause, [], [f960])).
fof(f960, plain, (spl45_81 <=> (e2 = op(e2, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl45_81])])).
fof(f5598, plain, (spl45_7 | ~ spl45_19 | ~ spl45_82), inference(avatar_split_clause, [], [f5593, f965, f593, f542])).
fof(f5593, plain, ((e2 = op(e3, e2)) | (~ spl45_19 | ~ spl45_82)), inference(backward_demodulation, [], [f967, f595])).
fof(f5597, plain, (~ spl45_3 | ~ spl45_19), inference(avatar_split_clause, [], [f5591, f593, f525])).
fof(f525, plain, (spl45_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_3])])).
fof(f5591, plain, (~ (e2 = op(e3, e3)) | ~ spl45_19), inference(backward_demodulation, [], [f217, f595])).
fof(f217, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5590, plain, (~ spl45_17 | ~ spl45_24 | spl45_149), inference(avatar_split_clause, [], [f5588, f1380, f614, f585])).
fof(f585, plain, (spl45_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_17])])).
fof(f1380, plain, (spl45_149 <=> (e0 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl45_149])])).
fof(f5588, plain, (~ (e0 = op(e2, e3)) | (~ spl45_24 | spl45_149)), inference(backward_demodulation, [], [f1382, f616])).
fof(f616, plain, ((e3 = op(e2, e2)) | ~ spl45_24), inference(avatar_component_clause, [], [f614])).
fof(f1382, plain, (~ (e0 = op(e2, op(e2, e2))) | spl45_149), inference(avatar_component_clause, [], [f1380])).
fof(f5589, plain, (~ spl45_8 | ~ spl45_24), inference(avatar_split_clause, [], [f5584, f614, f546])).
fof(f5584, plain, (~ (e3 = op(e3, e2)) | ~ spl45_24), inference(backward_demodulation, [], [f211, f616])).
fof(f211, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5583, plain, (~ spl45_17 | ~ spl45_29), inference(avatar_split_clause, [], [f5578, f636, f585])).
fof(f5578, plain, (~ (e0 = op(e2, e3)) | ~ spl45_29), inference(backward_demodulation, [], [f232, f638])).
fof(f638, plain, ((e0 = op(e2, e0)) | ~ spl45_29), inference(avatar_component_clause, [], [f636])).
fof(f232, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5560, plain, (~ spl45_36 | ~ spl45_44), inference(avatar_split_clause, [], [f5556, f699, f665])).
fof(f5556, plain, (~ (e3 = op(e1, e3)) | ~ spl45_44), inference(backward_demodulation, [], [f228, f701])).
fof(f701, plain, ((e3 = op(e1, e1)) | ~ spl45_44), inference(avatar_component_clause, [], [f699])).
fof(f228, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f5550, plain, (~ spl45_17 | ~ spl45_49), inference(avatar_split_clause, [], [f5543, f721, f585])).
fof(f5543, plain, (~ (e0 = op(e2, e3)) | ~ spl45_49), inference(backward_demodulation, [], [f213, f723])).
fof(f723, plain, ((e0 = op(e0, e3)) | ~ spl45_49), inference(avatar_component_clause, [], [f721])).
fof(f213, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5540, plain, (~ spl45_50 | ~ spl45_54), inference(avatar_split_clause, [], [f5534, f742, f725])).
fof(f742, plain, (spl45_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_54])])).
fof(f5534, plain, (~ (e1 = op(e0, e3)) | ~ spl45_54), inference(backward_demodulation, [], [f223, f744])).
fof(f744, plain, ((e1 = op(e0, e2)) | ~ spl45_54), inference(avatar_component_clause, [], [f742])).
fof(f223, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f5539, plain, (~ spl45_38 | ~ spl45_54), inference(avatar_split_clause, [], [f5531, f742, f674])).
fof(f5531, plain, (~ (e1 = op(e1, e2)) | ~ spl45_54), inference(backward_demodulation, [], [f206, f744])).
fof(f206, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f5529, plain, (~ spl45_38 | ~ spl45_59 | spl45_113), inference(avatar_split_clause, [], [f5522, f1114, f763, f674])).
fof(f5522, plain, (~ (e1 = op(e1, e2)) | (~ spl45_59 | spl45_113)), inference(backward_demodulation, [], [f1115, f765])).
fof(f765, plain, ((e2 = op(e0, e1)) | ~ spl45_59), inference(avatar_component_clause, [], [f763])).
fof(f5528, plain, (spl45_47 | ~ spl45_59 | ~ spl45_107), inference(avatar_split_clause, [], [f5521, f1085, f763, f712])).
fof(f5521, plain, ((e2 = op(e1, e0)) | (~ spl45_59 | ~ spl45_107)), inference(backward_demodulation, [], [f1087, f765])).
fof(f1087, plain, ((op(e0, e1) = op(e1, e0)) | ~ spl45_107), inference(avatar_component_clause, [], [f1085])).
fof(f5527, plain, (~ spl45_43 | ~ spl45_59), inference(avatar_split_clause, [], [f5517, f763, f695])).
fof(f695, plain, (spl45_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_43])])).
fof(f5517, plain, (~ (e2 = op(e1, e1)) | ~ spl45_59), inference(backward_demodulation, [], [f200, f765])).
fof(f5516, plain, (~ spl45_50 | ~ spl45_64 | spl45_137), inference(avatar_split_clause, [], [f5513, f1320, f784, f725])).
fof(f784, plain, (spl45_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl45_64])])).
fof(f1320, plain, (spl45_137 <=> (e1 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl45_137])])).
fof(f5513, plain, (~ (e1 = op(e0, e3)) | (~ spl45_64 | spl45_137)), inference(backward_demodulation, [], [f1322, f786])).
fof(f786, plain, ((op(e0, e0) = e3) | ~ spl45_64), inference(avatar_component_clause, [], [f784])).
fof(f1322, plain, (~ (e1 = op(e0, op(e0, e0))) | spl45_137), inference(avatar_component_clause, [], [f1320])).
fof(f5515, plain, (~ spl45_60 | ~ spl45_64), inference(avatar_split_clause, [], [f5507, f784, f767])).
fof(f5507, plain, (~ (e3 = op(e0, e1)) | ~ spl45_64), inference(backward_demodulation, [], [f218, f786])).
fof(f218, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f5514, plain, (~ spl45_32 | ~ spl45_64), inference(avatar_split_clause, [], [f5506, f784, f648])).
fof(f5506, plain, (~ (e3 = op(e2, e0)) | ~ spl45_64), inference(backward_demodulation, [], [f195, f786])).
fof(f5502, plain, (~ spl45_50 | ~ spl45_14 | spl45_88), inference(avatar_split_clause, [], [f5501, f993, f572, f725])).
fof(f5501, plain, (~ (e1 = op(e0, e3)) | (~ spl45_14 | spl45_88)), inference(forward_demodulation, [], [f994, f574])).
fof(f994, plain, (~ (op(e0, e3) = op(e3, e0)) | spl45_88), inference(avatar_component_clause, [], [f993])).
fof(f5500, plain, (~ spl45_38 | ~ spl45_26 | spl45_95), inference(avatar_split_clause, [], [f5499, f1027, f623, f674])).
fof(f5499, plain, (~ (e1 = op(e1, e2)) | (~ spl45_26 | spl45_95)), inference(forward_demodulation, [], [f1028, f625])).
fof(f1028, plain, (~ (op(e1, e2) = op(e2, e1)) | spl45_95), inference(avatar_component_clause, [], [f1027])).
fof(f5485, plain, (~ spl45_46 | ~ spl45_14), inference(avatar_split_clause, [], [f5484, f572, f708])).
fof(f708, plain, (spl45_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_46])])).
fof(f5484, plain, (~ (e1 = op(e1, e0)) | ~ spl45_14), inference(forward_demodulation, [], [f198, f574])).
fof(f198, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f5483, plain, (~ spl45_46 | ~ spl45_9 | spl45_84), inference(avatar_split_clause, [], [f5461, f974, f551, f708])).
fof(f974, plain, (spl45_84 <=> (e1 = op(e1, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl45_84])])).
fof(f5461, plain, (~ (e1 = op(e1, e0)) | (~ spl45_9 | spl45_84)), inference(backward_demodulation, [], [f975, f553])).
fof(f553, plain, ((e0 = op(e3, e1)) | ~ spl45_9), inference(avatar_component_clause, [], [f551])).
fof(f975, plain, (~ (e1 = op(e1, op(e3, e1))) | spl45_84), inference(avatar_component_clause, [], [f974])).
fof(f5477, plain, (~ spl45_18 | ~ spl45_26), inference(avatar_split_clause, [], [f5476, f623, f589])).
fof(f589, plain, (spl45_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_18])])).
fof(f5476, plain, (~ (e1 = op(e2, e3)) | ~ spl45_26), inference(forward_demodulation, [], [f234, f625])).
fof(f234, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5462, plain, (~ spl45_5 | ~ spl45_9), inference(avatar_split_clause, [], [f5456, f551, f534])).
fof(f534, plain, (spl45_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_5])])).
fof(f5456, plain, (~ (e0 = op(e3, e2)) | ~ spl45_9), inference(backward_demodulation, [], [f239, f553])).
fof(f5445, plain, (~ spl45_22 | ~ spl45_26), inference(avatar_split_clause, [], [f5441, f623, f606])).
fof(f606, plain, (spl45_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_22])])).
fof(f5441, plain, (~ (e1 = op(e2, e2)) | ~ spl45_26), inference(backward_demodulation, [], [f233, f625])).
fof(f233, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f5439, plain, (~ spl45_12 | ~ spl45_36 | spl45_85), inference(avatar_split_clause, [], [f5437, f979, f665, f563])).
fof(f979, plain, (spl45_85 <=> (op(e1, e3) = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_85])])).
fof(f5437, plain, (~ (e3 = op(e3, e1)) | (~ spl45_36 | spl45_85)), inference(backward_demodulation, [], [f980, f667])).
fof(f980, plain, (~ (op(e1, e3) = op(e3, e1)) | spl45_85), inference(avatar_component_clause, [], [f979])).
fof(f5434, plain, (~ spl45_22 | ~ spl45_38), inference(avatar_split_clause, [], [f5426, f674, f606])).
fof(f5426, plain, (~ (e1 = op(e2, e2)) | ~ spl45_38), inference(backward_demodulation, [], [f209, f676])).
fof(f209, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f5425, plain, (~ spl45_40 | ~ spl45_43 | spl45_139), inference(avatar_split_clause, [], [f5419, f1331, f695, f682])).
fof(f682, plain, (spl45_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_40])])).
fof(f1331, plain, (spl45_139 <=> (e3 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl45_139])])).
fof(f5419, plain, (~ (e3 = op(e1, e2)) | (~ spl45_43 | spl45_139)), inference(backward_demodulation, [], [f1333, f697])).
fof(f697, plain, ((e2 = op(e1, e1)) | ~ spl45_43), inference(avatar_component_clause, [], [f695])).
fof(f1333, plain, (~ (e3 = op(e1, op(e1, e1))) | spl45_139), inference(avatar_component_clause, [], [f1331])).
fof(f5424, plain, (~ spl45_35 | ~ spl45_43), inference(avatar_split_clause, [], [f5416, f695, f661])).
fof(f661, plain, (spl45_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_35])])).
fof(f5416, plain, (~ (e2 = op(e1, e3)) | ~ spl45_43), inference(backward_demodulation, [], [f228, f697])).
fof(f5423, plain, (~ spl45_27 | ~ spl45_43), inference(avatar_split_clause, [], [f5413, f695, f627])).
fof(f627, plain, (spl45_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_27])])).
fof(f5413, plain, (~ (e2 = op(e2, e1)) | ~ spl45_43), inference(backward_demodulation, [], [f203, f697])).
fof(f203, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f5411, plain, (spl45_14 | ~ spl45_50 | ~ spl45_88), inference(avatar_split_clause, [], [f5405, f993, f725, f572])).
fof(f5405, plain, ((e1 = op(e3, e0)) | (~ spl45_50 | ~ spl45_88)), inference(backward_demodulation, [], [f995, f727])).
fof(f5410, plain, (~ spl45_2 | ~ spl45_50), inference(avatar_split_clause, [], [f5402, f725, f521])).
fof(f521, plain, (spl45_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_2])])).
fof(f5402, plain, (~ (e1 = op(e3, e3)) | ~ spl45_50), inference(backward_demodulation, [], [f214, f727])).
fof(f214, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5400, plain, (spl45_31 | ~ spl45_53 | ~ spl45_111), inference(avatar_contradiction_clause, [], [f5399])).
fof(f5399, plain, ($false | (spl45_31 | ~ spl45_53 | ~ spl45_111)), inference(subsumption_resolution, [], [f5396, f645])).
fof(f645, plain, (~ (e2 = op(e2, e0)) | spl45_31), inference(avatar_component_clause, [], [f644])).
fof(f644, plain, (spl45_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_31])])).
fof(f5396, plain, ((e2 = op(e2, e0)) | (~ spl45_53 | ~ spl45_111)), inference(backward_demodulation, [], [f1106, f740])).
fof(f740, plain, ((e0 = op(e0, e2)) | ~ spl45_53), inference(avatar_component_clause, [], [f738])).
fof(f738, plain, (spl45_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_53])])).
fof(f1106, plain, ((e2 = op(e2, op(e0, e2))) | ~ spl45_111), inference(avatar_component_clause, [], [f1104])).
fof(f1104, plain, (spl45_111 <=> (e2 = op(e2, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl45_111])])).
fof(f5398, plain, (~ spl45_5 | ~ spl45_53), inference(avatar_split_clause, [], [f5391, f738, f534])).
fof(f5391, plain, (~ (e0 = op(e3, e2)) | ~ spl45_53), inference(backward_demodulation, [], [f208, f740])).
fof(f5388, plain, (~ spl45_52 | ~ spl45_60), inference(avatar_split_clause, [], [f5385, f767, f733])).
fof(f733, plain, (spl45_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_52])])).
fof(f5385, plain, (~ (e3 = op(e0, e3)) | ~ spl45_60), inference(backward_demodulation, [], [f222, f769])).
fof(f222, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f5387, plain, (~ spl45_12 | ~ spl45_60), inference(avatar_split_clause, [], [f5383, f767, f563])).
fof(f5383, plain, (~ (e3 = op(e3, e1)) | ~ spl45_60), inference(backward_demodulation, [], [f202, f769])).
fof(f202, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f5380, plain, (~ spl45_55 | ~ spl45_63), inference(avatar_split_clause, [], [f5373, f780, f746])).
fof(f780, plain, (spl45_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl45_63])])).
fof(f5373, plain, (~ (e2 = op(e0, e2)) | ~ spl45_63), inference(backward_demodulation, [], [f219, f782])).
fof(f782, plain, ((op(e0, e0) = e2) | ~ spl45_63), inference(avatar_component_clause, [], [f780])).
fof(f219, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f5379, plain, (~ spl45_15 | ~ spl45_63), inference(avatar_split_clause, [], [f5371, f780, f576])).
fof(f576, plain, (spl45_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_15])])).
fof(f5371, plain, (~ (e2 = op(e3, e0)) | ~ spl45_63), inference(backward_demodulation, [], [f196, f782])).
fof(f196, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f5364, plain, (~ spl45_32 | spl45_56 | ~ spl45_98), inference(avatar_contradiction_clause, [], [f5363])).
fof(f5363, plain, ($false | (~ spl45_32 | spl45_56 | ~ spl45_98)), inference(subsumption_resolution, [], [f5362, f751])).
fof(f751, plain, (~ (e3 = op(e0, e2)) | spl45_56), inference(avatar_component_clause, [], [f750])).
fof(f750, plain, (spl45_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_56])])).
fof(f5362, plain, ((e3 = op(e0, e2)) | (~ spl45_32 | ~ spl45_98)), inference(forward_demodulation, [], [f1043, f650])).
fof(f1043, plain, ((op(e0, e2) = op(e2, e0)) | ~ spl45_98), inference(avatar_component_clause, [], [f1041])).
fof(f1041, plain, (spl45_98 <=> (op(e0, e2) = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_98])])).
fof(f5361, plain, (~ spl45_57 | ~ spl45_45 | spl45_107), inference(avatar_split_clause, [], [f5360, f1085, f704, f755])).
fof(f755, plain, (spl45_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_57])])).
fof(f704, plain, (spl45_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_45])])).
fof(f5360, plain, (~ (e0 = op(e0, e1)) | (~ spl45_45 | spl45_107)), inference(forward_demodulation, [], [f1086, f706])).
fof(f706, plain, ((e0 = op(e1, e0)) | ~ spl45_45), inference(avatar_component_clause, [], [f704])).
fof(f5354, plain, (~ spl45_28 | ~ spl45_32), inference(avatar_split_clause, [], [f5353, f648, f631])).
fof(f631, plain, (spl45_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_28])])).
fof(f5353, plain, (~ (e3 = op(e2, e1)) | ~ spl45_32), inference(forward_demodulation, [], [f230, f650])).
fof(f230, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f5352, plain, (~ spl45_16 | ~ spl45_32), inference(avatar_split_clause, [], [f5351, f648, f580])).
fof(f580, plain, (spl45_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_16])])).
fof(f5351, plain, (~ (e3 = op(e3, e0)) | ~ spl45_32), inference(forward_demodulation, [], [f199, f650])).
fof(f199, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f5337, plain, (~ spl45_4 | ~ spl45_12), inference(avatar_split_clause, [], [f5336, f563, f529])).
fof(f5336, plain, (~ (e3 = op(e3, e3)) | ~ spl45_12), inference(backward_demodulation, [], [f240, f565])).
fof(f565, plain, ((e3 = op(e3, e1)) | ~ spl45_12), inference(avatar_component_clause, [], [f563])).
fof(f240, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5318, plain, (~ spl45_24 | ~ spl45_40), inference(avatar_split_clause, [], [f5316, f682, f614])).
fof(f5316, plain, (~ (e3 = op(e2, e2)) | ~ spl45_40), inference(backward_demodulation, [], [f209, f684])).
fof(f684, plain, ((e3 = op(e1, e2)) | ~ spl45_40), inference(avatar_component_clause, [], [f682])).
fof(f5312, plain, (~ spl45_42 | ~ spl45_138), inference(avatar_contradiction_clause, [], [f5311])).
fof(f5311, plain, ($false | (~ spl45_42 | ~ spl45_138)), inference(subsumption_resolution, [], [f5310, f245])).
fof(f245, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax4)).
fof(f5310, plain, ((e1 = e2) | (~ spl45_42 | ~ spl45_138)), inference(forward_demodulation, [], [f5309, f693])).
fof(f693, plain, ((e1 = op(e1, e1)) | ~ spl45_42), inference(avatar_component_clause, [], [f691])).
fof(f691, plain, (spl45_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_42])])).
fof(f5309, plain, ((e2 = op(e1, e1)) | (~ spl45_42 | ~ spl45_138)), inference(forward_demodulation, [], [f5303, f693])).
fof(f5303, plain, ((e2 = op(op(e1, e1), e1)) | (~ spl45_42 | ~ spl45_138)), inference(backward_demodulation, [], [f1328, f693])).
fof(f1328, plain, ((e2 = op(op(e1, op(e1, e1)), e1)) | ~ spl45_138), inference(avatar_component_clause, [], [f1327])).
fof(f1327, plain, (spl45_138 <=> (e2 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_138])])).
fof(f5308, plain, (~ spl45_38 | ~ spl45_42), inference(avatar_split_clause, [], [f5300, f691, f674])).
fof(f5300, plain, (~ (e1 = op(e1, e2)) | ~ spl45_42), inference(backward_demodulation, [], [f227, f693])).
fof(f5307, plain, (~ spl45_10 | ~ spl45_42), inference(avatar_split_clause, [], [f5299, f691, f555])).
fof(f5299, plain, (~ (e1 = op(e3, e1)) | ~ spl45_42), inference(backward_demodulation, [], [f204, f693])).
fof(f5297, plain, (~ spl45_4 | ~ spl45_52), inference(avatar_split_clause, [], [f5295, f733, f529])).
fof(f5295, plain, (~ (e3 = op(e3, e3)) | ~ spl45_52), inference(backward_demodulation, [], [f214, f735])).
fof(f735, plain, ((e3 = op(e0, e3)) | ~ spl45_52), inference(avatar_component_clause, [], [f733])).
fof(f5284, plain, (~ spl45_50 | ~ spl45_62), inference(avatar_split_clause, [], [f5274, f776, f725])).
fof(f776, plain, (spl45_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl45_62])])).
fof(f5274, plain, (~ (e1 = op(e0, e3)) | ~ spl45_62), inference(backward_demodulation, [], [f220, f778])).
fof(f778, plain, ((op(e0, e0) = e1) | ~ spl45_62), inference(avatar_component_clause, [], [f776])).
fof(f5283, plain, (~ spl45_30 | ~ spl45_62), inference(avatar_split_clause, [], [f5273, f776, f640])).
fof(f640, plain, (spl45_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_30])])).
fof(f5273, plain, (~ (e1 = op(e2, e0)) | ~ spl45_62), inference(backward_demodulation, [], [f195, f778])).
fof(f5257, plain, (~ spl45_39 | ~ spl45_35), inference(avatar_split_clause, [], [f5256, f661, f678])).
fof(f678, plain, (spl45_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_39])])).
fof(f5256, plain, (~ (e2 = op(e1, e2)) | ~ spl45_35), inference(forward_demodulation, [], [f229, f663])).
fof(f663, plain, ((e2 = op(e1, e3)) | ~ spl45_35), inference(avatar_component_clause, [], [f661])).
fof(f5254, plain, (spl45_31 | ~ spl45_5 | ~ spl45_81), inference(avatar_split_clause, [], [f5240, f960, f534, f644])).
fof(f5240, plain, ((e2 = op(e2, e0)) | (~ spl45_5 | ~ spl45_81)), inference(forward_demodulation, [], [f962, f536])).
fof(f536, plain, ((e0 = op(e3, e2)) | ~ spl45_5), inference(avatar_component_clause, [], [f534])).
fof(f5249, plain, (~ spl45_11 | ~ spl45_35 | spl45_85), inference(avatar_split_clause, [], [f5248, f979, f661, f559])).
fof(f5248, plain, (~ (e2 = op(e3, e1)) | (~ spl45_35 | spl45_85)), inference(forward_demodulation, [], [f980, f663])).
fof(f5232, plain, (~ spl45_15 | ~ spl45_17 | ~ spl45_90), inference(avatar_contradiction_clause, [], [f5231])).
fof(f5231, plain, ($false | (~ spl45_15 | ~ spl45_17 | ~ spl45_90)), inference(subsumption_resolution, [], [f5230, f247])).
fof(f247, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f5230, plain, ((e2 = e3) | (~ spl45_15 | ~ spl45_17 | ~ spl45_90)), inference(forward_demodulation, [], [f5227, f578])).
fof(f578, plain, ((e2 = op(e3, e0)) | ~ spl45_15), inference(avatar_component_clause, [], [f576])).
fof(f5227, plain, ((e3 = op(e3, e0)) | (~ spl45_17 | ~ spl45_90)), inference(backward_demodulation, [], [f1004, f587])).
fof(f587, plain, ((e0 = op(e2, e3)) | ~ spl45_17), inference(avatar_component_clause, [], [f585])).
fof(f5223, plain, (~ spl45_30 | ~ spl45_57 | spl45_97), inference(avatar_contradiction_clause, [], [f5222])).
fof(f5222, plain, ($false | (~ spl45_30 | ~ spl45_57 | spl45_97)), inference(subsumption_resolution, [], [f5220, f757])).
fof(f757, plain, ((e0 = op(e0, e1)) | ~ spl45_57), inference(avatar_component_clause, [], [f755])).
fof(f5220, plain, (~ (e0 = op(e0, e1)) | (~ spl45_30 | spl45_97)), inference(backward_demodulation, [], [f1037, f642])).
fof(f642, plain, ((e1 = op(e2, e0)) | ~ spl45_30), inference(avatar_component_clause, [], [f640])).
fof(f1037, plain, (~ (e0 = op(e0, op(e2, e0))) | spl45_97), inference(avatar_component_clause, [], [f1036])).
fof(f1036, plain, (spl45_97 <=> (e0 = op(e0, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl45_97])])).
fof(f5221, plain, (~ spl45_18 | ~ spl45_30), inference(avatar_split_clause, [], [f5217, f640, f589])).
fof(f5217, plain, (~ (e1 = op(e2, e3)) | ~ spl45_30), inference(backward_demodulation, [], [f232, f642])).
fof(f5206, plain, (~ spl45_33 | ~ spl45_45), inference(avatar_split_clause, [], [f5200, f704, f653])).
fof(f653, plain, (spl45_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_33])])).
fof(f5200, plain, (~ (e0 = op(e1, e3)) | ~ spl45_45), inference(backward_demodulation, [], [f226, f706])).
fof(f226, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f5205, plain, (~ spl45_29 | ~ spl45_45), inference(avatar_split_clause, [], [f5198, f704, f636])).
fof(f5198, plain, (~ (e0 = op(e2, e0)) | ~ spl45_45), inference(backward_demodulation, [], [f197, f706])).
fof(f197, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f5195, plain, (~ spl45_18 | ~ spl45_50), inference(avatar_split_clause, [], [f5188, f725, f589])).
fof(f5188, plain, (~ (e1 = op(e2, e3)) | ~ spl45_50), inference(backward_demodulation, [], [f213, f727])).
fof(f5186, plain, (~ spl45_51 | ~ spl45_55), inference(avatar_split_clause, [], [f5183, f746, f729])).
fof(f729, plain, (spl45_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_51])])).
fof(f5183, plain, (~ (e2 = op(e0, e3)) | ~ spl45_55), inference(backward_demodulation, [], [f223, f748])).
fof(f5185, plain, (~ spl45_39 | ~ spl45_55), inference(avatar_split_clause, [], [f5181, f746, f678])).
fof(f5181, plain, (~ (e2 = op(e1, e2)) | ~ spl45_55), inference(backward_demodulation, [], [f206, f748])).
fof(f5179, plain, (spl45_17 | ~ spl45_5 | ~ spl45_82), inference(avatar_split_clause, [], [f5178, f965, f534, f585])).
fof(f5178, plain, ((e0 = op(e2, e3)) | (~ spl45_5 | ~ spl45_82)), inference(forward_demodulation, [], [f967, f536])).
fof(f5174, plain, (spl45_45 | ~ spl45_57 | ~ spl45_107), inference(avatar_split_clause, [], [f5173, f1085, f755, f704])).
fof(f5173, plain, ((e0 = op(e1, e0)) | (~ spl45_57 | ~ spl45_107)), inference(forward_demodulation, [], [f1087, f757])).
fof(f5172, plain, (~ spl45_46 | ~ spl45_57 | spl45_113), inference(avatar_split_clause, [], [f5171, f1114, f755, f708])).
fof(f5171, plain, (~ (e1 = op(e1, e0)) | (~ spl45_57 | spl45_113)), inference(forward_demodulation, [], [f1115, f757])).
fof(f5169, plain, (~ spl45_51 | ~ spl45_64 | spl45_135), inference(avatar_split_clause, [], [f5168, f1310, f784, f729])).
fof(f1310, plain, (spl45_135 <=> (e2 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl45_135])])).
fof(f5168, plain, (~ (e2 = op(e0, e3)) | (~ spl45_64 | spl45_135)), inference(forward_demodulation, [], [f1312, f786])).
fof(f1312, plain, (~ (e2 = op(e0, op(e0, e0))) | spl45_135), inference(avatar_component_clause, [], [f1310])).
fof(f5161, plain, (~ spl45_33 | ~ spl45_44 | spl45_143), inference(avatar_split_clause, [], [f5160, f1350, f699, f653])).
fof(f1350, plain, (spl45_143 <=> (e0 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl45_143])])).
fof(f5160, plain, (~ (e0 = op(e1, e3)) | (~ spl45_44 | spl45_143)), inference(forward_demodulation, [], [f1352, f701])).
fof(f1352, plain, (~ (e0 = op(e1, op(e1, e1))) | spl45_143), inference(avatar_component_clause, [], [f1350])).
fof(f5159, plain, (~ spl45_18 | ~ spl45_24 | spl45_147), inference(avatar_split_clause, [], [f5158, f1370, f614, f589])).
fof(f1370, plain, (spl45_147 <=> (e1 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl45_147])])).
fof(f5158, plain, (~ (e1 = op(e2, e3)) | (~ spl45_24 | spl45_147)), inference(forward_demodulation, [], [f1372, f616])).
fof(f1372, plain, (~ (e1 = op(e2, op(e2, e2))) | spl45_147), inference(avatar_component_clause, [], [f1370])).
fof(f5143, plain, (~ spl45_47 | ~ spl45_15), inference(avatar_split_clause, [], [f5142, f576, f712])).
fof(f5142, plain, (~ (e2 = op(e1, e0)) | ~ spl45_15), inference(forward_demodulation, [], [f198, f578])).
fof(f5139, plain, (~ spl45_34 | ~ spl45_10 | spl45_85), inference(avatar_split_clause, [], [f5138, f979, f555, f657])).
fof(f5138, plain, (~ (e1 = op(e1, e3)) | (~ spl45_10 | spl45_85)), inference(forward_demodulation, [], [f980, f557])).
fof(f557, plain, ((e1 = op(e3, e1)) | ~ spl45_10), inference(avatar_component_clause, [], [f555])).
fof(f5133, plain, (~ spl45_31 | ~ spl45_15), inference(avatar_split_clause, [], [f5132, f576, f644])).
fof(f5132, plain, (~ (e2 = op(e2, e0)) | ~ spl45_15), inference(forward_demodulation, [], [f199, f578])).
fof(f5131, plain, (~ spl45_31 | ~ spl45_5 | spl45_81), inference(avatar_split_clause, [], [f5128, f960, f534, f644])).
fof(f5128, plain, (~ (e2 = op(e2, e0)) | (~ spl45_5 | spl45_81)), inference(backward_demodulation, [], [f961, f536])).
fof(f961, plain, (~ (e2 = op(e2, op(e3, e2))) | spl45_81), inference(avatar_component_clause, [], [f960])).
fof(f5120, plain, (~ spl45_15 | spl45_53 | ~ spl45_87), inference(avatar_contradiction_clause, [], [f5119])).
fof(f5119, plain, ($false | (~ spl45_15 | spl45_53 | ~ spl45_87)), inference(subsumption_resolution, [], [f5117, f739])).
fof(f739, plain, (~ (e0 = op(e0, e2)) | spl45_53), inference(avatar_component_clause, [], [f738])).
fof(f5117, plain, ((e0 = op(e0, e2)) | (~ spl45_15 | ~ spl45_87)), inference(backward_demodulation, [], [f990, f578])).
fof(f990, plain, ((e0 = op(e0, op(e3, e0))) | ~ spl45_87), inference(avatar_component_clause, [], [f988])).
fof(f988, plain, (spl45_87 <=> (e0 = op(e0, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl45_87])])).
fof(f5118, plain, (~ spl45_7 | ~ spl45_15), inference(avatar_split_clause, [], [f5115, f576, f542])).
fof(f5115, plain, (~ (e2 = op(e3, e2)) | ~ spl45_15), inference(backward_demodulation, [], [f237, f578])).
fof(f237, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5113, plain, (~ spl45_7 | ~ spl45_39), inference(avatar_split_clause, [], [f5112, f678, f542])).
fof(f5112, plain, (~ (e2 = op(e3, e2)) | ~ spl45_39), inference(backward_demodulation, [], [f210, f680])).
fof(f680, plain, ((e2 = op(e1, e2)) | ~ spl45_39), inference(avatar_component_clause, [], [f678])).
fof(f210, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5109, plain, (~ spl45_38 | ~ spl45_46), inference(avatar_split_clause, [], [f5107, f708, f674])).
fof(f5107, plain, (~ (e1 = op(e1, e2)) | ~ spl45_46), inference(backward_demodulation, [], [f225, f710])).
fof(f710, plain, ((e1 = op(e1, e0)) | ~ spl45_46), inference(avatar_component_clause, [], [f708])).
fof(f225, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f5096, plain, (spl45_46 | ~ spl45_57 | ~ spl45_113), inference(avatar_split_clause, [], [f5093, f1114, f755, f708])).
fof(f5093, plain, ((e1 = op(e1, e0)) | (~ spl45_57 | ~ spl45_113)), inference(backward_demodulation, [], [f1116, f757])).
fof(f1116, plain, ((e1 = op(e1, op(e0, e1))) | ~ spl45_113), inference(avatar_component_clause, [], [f1114])).
fof(f5095, plain, (~ spl45_53 | ~ spl45_57), inference(avatar_split_clause, [], [f5088, f755, f738])).
fof(f5088, plain, (~ (e0 = op(e0, e2)) | ~ spl45_57), inference(backward_demodulation, [], [f221, f757])).
fof(f221, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f5094, plain, (~ spl45_9 | ~ spl45_57), inference(avatar_split_clause, [], [f5087, f755, f551])).
fof(f5087, plain, (~ (e0 = op(e3, e1)) | ~ spl45_57), inference(backward_demodulation, [], [f202, f757])).
fof(f5084, plain, (spl45_15 | ~ spl45_51 | ~ spl45_88), inference(avatar_split_clause, [], [f5083, f993, f729, f576])).
fof(f5083, plain, ((e2 = op(e3, e0)) | (~ spl45_51 | ~ spl45_88)), inference(forward_demodulation, [], [f995, f731])).
fof(f731, plain, ((e2 = op(e0, e3)) | ~ spl45_51), inference(avatar_component_clause, [], [f729])).
fof(f5082, plain, (~ spl45_53 | ~ spl45_29 | spl45_98), inference(avatar_split_clause, [], [f5081, f1041, f636, f738])).
fof(f5081, plain, (~ (e0 = op(e0, e2)) | (~ spl45_29 | spl45_98)), inference(forward_demodulation, [], [f1042, f638])).
fof(f1042, plain, (~ (op(e0, e2) = op(e2, e0)) | spl45_98), inference(avatar_component_clause, [], [f1041])).
fof(f5068, plain, (~ spl45_48 | ~ spl45_64), inference(avatar_split_clause, [], [f5067, f784, f716])).
fof(f716, plain, (spl45_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_48])])).
fof(f5067, plain, (~ (e3 = op(e1, e0)) | ~ spl45_64), inference(forward_demodulation, [], [f194, f786])).
fof(f194, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f5066, plain, (~ spl45_48 | ~ spl45_44), inference(avatar_split_clause, [], [f5065, f699, f716])).
fof(f5065, plain, (~ (e3 = op(e1, e0)) | ~ spl45_44), inference(forward_demodulation, [], [f224, f701])).
fof(f224, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f5062, plain, (~ spl45_39 | ~ spl45_27 | spl45_95), inference(avatar_split_clause, [], [f5061, f1027, f627, f678])).
fof(f5061, plain, (~ (e2 = op(e1, e2)) | (~ spl45_27 | spl45_95)), inference(forward_demodulation, [], [f1028, f629])).
fof(f629, plain, ((e2 = op(e2, e1)) | ~ spl45_27), inference(avatar_component_clause, [], [f627])).
fof(f5060, plain, (spl45_38 | ~ spl45_27 | ~ spl45_94), inference(avatar_split_clause, [], [f5044, f1022, f627, f674])).
fof(f1022, plain, (spl45_94 <=> (e1 = op(e1, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl45_94])])).
fof(f5044, plain, ((e1 = op(e1, e2)) | (~ spl45_27 | ~ spl45_94)), inference(forward_demodulation, [], [f1024, f629])).
fof(f1024, plain, ((e1 = op(e1, op(e2, e1))) | ~ spl45_94), inference(avatar_component_clause, [], [f1022])).
fof(f5058, plain, (spl45_16 | ~ spl45_33 | ~ spl45_100), inference(avatar_split_clause, [], [f5039, f1050, f653, f580])).
fof(f1050, plain, (spl45_100 <=> (e3 = op(e3, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl45_100])])).
fof(f5039, plain, ((e3 = op(e3, e0)) | (~ spl45_33 | ~ spl45_100)), inference(forward_demodulation, [], [f1052, f655])).
fof(f655, plain, ((e0 = op(e1, e3)) | ~ spl45_33), inference(avatar_component_clause, [], [f653])).
fof(f1052, plain, ((e3 = op(e3, op(e1, e3))) | ~ spl45_100), inference(avatar_component_clause, [], [f1050])).
fof(f5055, plain, (~ spl45_11 | ~ spl45_27), inference(avatar_split_clause, [], [f5054, f627, f559])).
fof(f5054, plain, (~ (e2 = op(e3, e1)) | ~ spl45_27), inference(forward_demodulation, [], [f205, f629])).
fof(f205, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f5053, plain, (~ spl45_12 | ~ spl45_44), inference(avatar_split_clause, [], [f5052, f699, f563])).
fof(f5052, plain, (~ (e3 = op(e3, e1)) | ~ spl45_44), inference(forward_demodulation, [], [f204, f701])).
fof(f5051, plain, (~ spl45_12 | ~ spl45_18 | spl45_90), inference(avatar_split_clause, [], [f5019, f1002, f589, f563])).
fof(f5019, plain, (~ (e3 = op(e3, e1)) | (~ spl45_18 | spl45_90)), inference(backward_demodulation, [], [f1003, f591])).
fof(f591, plain, ((e1 = op(e2, e3)) | ~ spl45_18), inference(avatar_component_clause, [], [f589])).
fof(f1003, plain, (~ (e3 = op(e3, op(e2, e3))) | spl45_90), inference(avatar_component_clause, [], [f1002])).
fof(f5026, plain, (~ spl45_2 | ~ spl45_14), inference(avatar_split_clause, [], [f5022, f572, f521])).
fof(f5022, plain, (~ (e1 = op(e3, e3)) | ~ spl45_14), inference(backward_demodulation, [], [f238, f574])).
fof(f238, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5020, plain, (~ spl45_2 | ~ spl45_18), inference(avatar_split_clause, [], [f5016, f589, f521])).
fof(f5016, plain, (~ (e1 = op(e3, e3)) | ~ spl45_18), inference(backward_demodulation, [], [f217, f591])).
fof(f5009, plain, (~ spl45_19 | ~ spl45_27), inference(avatar_split_clause, [], [f5006, f627, f593])).
fof(f5006, plain, (~ (e2 = op(e2, e3)) | ~ spl45_27), inference(backward_demodulation, [], [f234, f629])).
fof(f5004, plain, (~ spl45_16 | ~ spl45_33 | spl45_100), inference(avatar_split_clause, [], [f5003, f1050, f653, f580])).
fof(f5003, plain, (~ (e3 = op(e3, e0)) | (~ spl45_33 | spl45_100)), inference(backward_demodulation, [], [f1051, f655])).
fof(f1051, plain, (~ (e3 = op(e3, op(e1, e3))) | spl45_100), inference(avatar_component_clause, [], [f1050])).
fof(f4987, plain, (~ spl45_28 | ~ spl45_44), inference(avatar_split_clause, [], [f4983, f699, f631])).
fof(f4983, plain, (~ (e3 = op(e2, e1)) | ~ spl45_44), inference(backward_demodulation, [], [f203, f701])).
fof(f4981, plain, (~ spl45_19 | ~ spl45_51), inference(avatar_split_clause, [], [f4976, f729, f593])).
fof(f4976, plain, (~ (e2 = op(e2, e3)) | ~ spl45_51), inference(backward_demodulation, [], [f213, f731])).
fof(f4973, plain, (~ spl45_49 | ~ spl45_53), inference(avatar_split_clause, [], [f4967, f738, f721])).
fof(f4967, plain, (~ (e0 = op(e0, e3)) | ~ spl45_53), inference(backward_demodulation, [], [f223, f740])).
fof(f4972, plain, (~ spl45_37 | ~ spl45_53), inference(avatar_split_clause, [], [f4965, f738, f670])).
fof(f4965, plain, (~ (e0 = op(e1, e2)) | ~ spl45_53), inference(backward_demodulation, [], [f206, f740])).
fof(f4963, plain, (~ spl45_42 | ~ spl45_58), inference(avatar_split_clause, [], [f4956, f759, f691])).
fof(f759, plain, (spl45_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_58])])).
fof(f4956, plain, (~ (e1 = op(e1, e1)) | ~ spl45_58), inference(backward_demodulation, [], [f200, f761])).
fof(f761, plain, ((e1 = op(e0, e1)) | ~ spl45_58), inference(avatar_component_clause, [], [f759])).
fof(f4955, plain, (spl45_51 | ~ spl45_64 | ~ spl45_135), inference(avatar_split_clause, [], [f4951, f1310, f784, f729])).
fof(f4951, plain, ((e2 = op(e0, e3)) | (~ spl45_64 | ~ spl45_135)), inference(backward_demodulation, [], [f1311, f786])).
fof(f1311, plain, ((e2 = op(e0, op(e0, e0))) | ~ spl45_135), inference(avatar_component_clause, [], [f1310])).
fof(f4954, plain, (~ spl45_56 | ~ spl45_64), inference(avatar_split_clause, [], [f4949, f784, f750])).
fof(f4949, plain, (~ (e3 = op(e0, e2)) | ~ spl45_64), inference(backward_demodulation, [], [f219, f786])).
fof(f4953, plain, (~ spl45_16 | ~ spl45_64), inference(avatar_split_clause, [], [f4947, f784, f580])).
fof(f4947, plain, (~ (e3 = op(e3, e0)) | ~ spl45_64), inference(backward_demodulation, [], [f196, f786])).
fof(f4944, plain, (spl45_33 | ~ spl45_9 | ~ spl45_85), inference(avatar_split_clause, [], [f4943, f979, f551, f653])).
fof(f4943, plain, ((e0 = op(e1, e3)) | (~ spl45_9 | ~ spl45_85)), inference(forward_demodulation, [], [f981, f553])).
fof(f981, plain, ((op(e1, e3) = op(e3, e1)) | ~ spl45_85), inference(avatar_component_clause, [], [f979])).
fof(f4942, plain, (~ spl45_29 | spl45_61 | ~ spl45_97), inference(avatar_contradiction_clause, [], [f4941])).
fof(f4941, plain, ($false | (~ spl45_29 | spl45_61 | ~ spl45_97)), inference(subsumption_resolution, [], [f4940, f773])).
fof(f773, plain, (~ (e0 = op(e0, e0)) | spl45_61), inference(avatar_component_clause, [], [f772])).
fof(f4940, plain, ((e0 = op(e0, e0)) | (~ spl45_29 | ~ spl45_97)), inference(forward_demodulation, [], [f1038, f638])).
fof(f1038, plain, ((e0 = op(e0, op(e2, e0))) | ~ spl45_97), inference(avatar_component_clause, [], [f1036])).
fof(f4939, plain, (spl45_53 | ~ spl45_29 | ~ spl45_98), inference(avatar_split_clause, [], [f4938, f1041, f636, f738])).
fof(f4938, plain, ((e0 = op(e0, e2)) | (~ spl45_29 | ~ spl45_98)), inference(forward_demodulation, [], [f1043, f638])).
fof(f4935, plain, (~ spl45_59 | spl45_138 | ~ spl45_143), inference(avatar_split_clause, [], [f4934, f1350, f1327, f763])).
fof(f4934, plain, (~ (e2 = op(e0, e1)) | (spl45_138 | ~ spl45_143)), inference(forward_demodulation, [], [f1329, f1351])).
fof(f1351, plain, ((e0 = op(e1, op(e1, e1))) | ~ spl45_143), inference(avatar_component_clause, [], [f1350])).
fof(f1329, plain, (~ (e2 = op(op(e1, op(e1, e1)), e1)) | spl45_138), inference(avatar_component_clause, [], [f1327])).
fof(f4927, plain, (~ spl45_37 | ~ spl45_147 | spl45_148), inference(avatar_split_clause, [], [f4925, f1375, f1370, f670])).
fof(f1375, plain, (spl45_148 <=> (e0 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_148])])).
fof(f4925, plain, (~ (e0 = op(e1, e2)) | (~ spl45_147 | spl45_148)), inference(backward_demodulation, [], [f1377, f1371])).
fof(f1371, plain, ((e1 = op(e2, op(e2, e2))) | ~ spl45_147), inference(avatar_component_clause, [], [f1370])).
fof(f1377, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | spl45_148), inference(avatar_component_clause, [], [f1375])).
fof(f4919, plain, (~ spl45_39 | ~ spl45_47), inference(avatar_split_clause, [], [f4918, f712, f678])).
fof(f4918, plain, (~ (e2 = op(e1, e2)) | ~ spl45_47), inference(forward_demodulation, [], [f225, f714])).
fof(f4905, plain, (spl45_53 | ~ spl45_59 | ~ spl45_106 | ~ spl45_107), inference(avatar_contradiction_clause, [], [f4904])).
fof(f4904, plain, ($false | (spl45_53 | ~ spl45_59 | ~ spl45_106 | ~ spl45_107)), inference(subsumption_resolution, [], [f4900, f739])).
fof(f4900, plain, ((e0 = op(e0, e2)) | (~ spl45_59 | ~ spl45_106 | ~ spl45_107)), inference(backward_demodulation, [], [f4881, f765])).
fof(f4881, plain, ((e0 = op(e0, op(e0, e1))) | (~ spl45_106 | ~ spl45_107)), inference(forward_demodulation, [], [f1082, f1087])).
fof(f1082, plain, ((e0 = op(e0, op(e1, e0))) | ~ spl45_106), inference(avatar_component_clause, [], [f1080])).
fof(f1080, plain, (spl45_106 <=> (e0 = op(e0, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl45_106])])).
fof(f4893, plain, (~ spl45_58 | ~ spl45_62), inference(avatar_split_clause, [], [f4887, f776, f759])).
fof(f4887, plain, (~ (e1 = op(e0, e1)) | ~ spl45_62), inference(backward_demodulation, [], [f218, f778])).
fof(f4892, plain, (~ spl45_46 | ~ spl45_62), inference(avatar_split_clause, [], [f4886, f776, f708])).
fof(f4886, plain, (~ (e1 = op(e1, e0)) | ~ spl45_62), inference(backward_demodulation, [], [f194, f778])).
fof(f4873, plain, (spl45_59 | ~ spl45_138 | ~ spl45_143), inference(avatar_split_clause, [], [f4872, f1350, f1327, f763])).
fof(f4872, plain, ((e2 = op(e0, e1)) | (~ spl45_138 | ~ spl45_143)), inference(forward_demodulation, [], [f1328, f1351])).
fof(f4853, plain, (~ spl45_7 | ~ spl45_22 | ~ spl45_81), inference(avatar_contradiction_clause, [], [f4852])).
fof(f4852, plain, ($false | (~ spl45_7 | ~ spl45_22 | ~ spl45_81)), inference(subsumption_resolution, [], [f4851, f245])).
fof(f4851, plain, ((e1 = e2) | (~ spl45_7 | ~ spl45_22 | ~ spl45_81)), inference(forward_demodulation, [], [f4850, f608])).
fof(f608, plain, ((e1 = op(e2, e2)) | ~ spl45_22), inference(avatar_component_clause, [], [f606])).
fof(f4850, plain, ((e2 = op(e2, e2)) | (~ spl45_7 | ~ spl45_81)), inference(forward_demodulation, [], [f962, f544])).
fof(f544, plain, ((e2 = op(e3, e2)) | ~ spl45_7), inference(avatar_component_clause, [], [f542])).
fof(f4829, plain, (~ spl45_6 | ~ spl45_22), inference(avatar_split_clause, [], [f4822, f606, f538])).
fof(f538, plain, (spl45_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_6])])).
fof(f4822, plain, (~ (e1 = op(e3, e2)) | ~ spl45_22), inference(backward_demodulation, [], [f211, f608])).
fof(f4821, plain, (~ spl45_20 | ~ spl45_28), inference(avatar_split_clause, [], [f4819, f631, f597])).
fof(f597, plain, (spl45_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_20])])).
fof(f4819, plain, (~ (e3 = op(e2, e3)) | ~ spl45_28), inference(backward_demodulation, [], [f234, f633])).
fof(f633, plain, ((e3 = op(e2, e1)) | ~ spl45_28), inference(avatar_component_clause, [], [f631])).
fof(f4817, plain, (~ spl45_21 | ~ spl45_29), inference(avatar_split_clause, [], [f4811, f636, f602])).
fof(f602, plain, (spl45_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_21])])).
fof(f4811, plain, (~ (e0 = op(e2, e2)) | ~ spl45_29), inference(backward_demodulation, [], [f231, f638])).
fof(f4809, plain, (spl45_4 | ~ spl45_36 | ~ spl45_100), inference(avatar_contradiction_clause, [], [f4808])).
fof(f4808, plain, ($false | (spl45_4 | ~ spl45_36 | ~ spl45_100)), inference(subsumption_resolution, [], [f4806, f530])).
fof(f530, plain, (~ (e3 = op(e3, e3)) | spl45_4), inference(avatar_component_clause, [], [f529])).
fof(f4806, plain, ((e3 = op(e3, e3)) | (~ spl45_36 | ~ spl45_100)), inference(backward_demodulation, [], [f1052, f667])).
fof(f4807, plain, (~ spl45_20 | ~ spl45_36), inference(avatar_split_clause, [], [f4804, f665, f597])).
fof(f4804, plain, (~ (e3 = op(e2, e3)) | ~ spl45_36), inference(backward_demodulation, [], [f215, f667])).
fof(f215, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4802, plain, (~ spl45_31 | ~ spl45_37 | spl45_102), inference(avatar_split_clause, [], [f4797, f1060, f670, f644])).
fof(f1060, plain, (spl45_102 <=> (e2 = op(e2, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl45_102])])).
fof(f4797, plain, (~ (e2 = op(e2, e0)) | (~ spl45_37 | spl45_102)), inference(backward_demodulation, [], [f1061, f672])).
fof(f672, plain, ((e0 = op(e1, e2)) | ~ spl45_37), inference(avatar_component_clause, [], [f670])).
fof(f1061, plain, (~ (e2 = op(e2, op(e1, e2))) | spl45_102), inference(avatar_component_clause, [], [f1060])).
fof(f4801, plain, (~ spl45_33 | ~ spl45_37), inference(avatar_split_clause, [], [f4792, f670, f653])).
fof(f4792, plain, (~ (e0 = op(e1, e3)) | ~ spl45_37), inference(backward_demodulation, [], [f229, f672])).
fof(f4800, plain, (~ spl45_21 | ~ spl45_37), inference(avatar_split_clause, [], [f4790, f670, f602])).
fof(f4790, plain, (~ (e0 = op(e2, e2)) | ~ spl45_37), inference(backward_demodulation, [], [f209, f672])).
fof(f4789, plain, (~ spl45_33 | ~ spl45_49), inference(avatar_split_clause, [], [f4782, f721, f653])).
fof(f4782, plain, (~ (e0 = op(e1, e3)) | ~ spl45_49), inference(backward_demodulation, [], [f212, f723])).
fof(f4781, plain, (spl45_19 | ~ spl45_56 | ~ spl45_111), inference(avatar_split_clause, [], [f4779, f1104, f750, f593])).
fof(f4779, plain, ((e2 = op(e2, e3)) | (~ spl45_56 | ~ spl45_111)), inference(backward_demodulation, [], [f1106, f752])).
fof(f752, plain, ((e3 = op(e0, e2)) | ~ spl45_56), inference(avatar_component_clause, [], [f750])).
fof(f4780, plain, (~ spl45_40 | ~ spl45_56), inference(avatar_split_clause, [], [f4774, f750, f682])).
fof(f4774, plain, (~ (e3 = op(e1, e2)) | ~ spl45_56), inference(backward_demodulation, [], [f206, f752])).
fof(f4773, plain, (~ spl45_50 | ~ spl45_58), inference(avatar_split_clause, [], [f4768, f759, f725])).
fof(f4768, plain, (~ (e1 = op(e0, e3)) | ~ spl45_58), inference(backward_demodulation, [], [f222, f761])).
fof(f4772, plain, (~ spl45_26 | ~ spl45_58), inference(avatar_split_clause, [], [f4766, f759, f623])).
fof(f4766, plain, (~ (e1 = op(e2, e1)) | ~ spl45_58), inference(backward_demodulation, [], [f201, f761])).
fof(f201, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4762, plain, (~ spl45_31 | ~ spl45_63), inference(avatar_split_clause, [], [f4754, f780, f644])).
fof(f4754, plain, (~ (e2 = op(e2, e0)) | ~ spl45_63), inference(backward_demodulation, [], [f195, f782])).
fof(f4749, plain, (spl45_49 | ~ spl45_16 | ~ spl45_87), inference(avatar_split_clause, [], [f4748, f988, f580, f721])).
fof(f4748, plain, ((e0 = op(e0, e3)) | (~ spl45_16 | ~ spl45_87)), inference(forward_demodulation, [], [f990, f582])).
fof(f582, plain, ((e3 = op(e3, e0)) | ~ spl45_16), inference(avatar_component_clause, [], [f580])).
fof(f4747, plain, (spl45_58 | ~ spl45_46 | ~ spl45_107), inference(avatar_split_clause, [], [f4746, f1085, f708, f759])).
fof(f4746, plain, ((e1 = op(e0, e1)) | (~ spl45_46 | ~ spl45_107)), inference(forward_demodulation, [], [f1087, f710])).
fof(f4730, plain, (~ spl45_50 | spl45_150 | ~ spl45_155), inference(avatar_split_clause, [], [f4728, f1410, f1387, f725])).
fof(f1387, plain, (spl45_150 <=> (e1 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_150])])).
fof(f1410, plain, (spl45_155 <=> (e0 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl45_155])])).
fof(f4728, plain, (~ (e1 = op(e0, e3)) | (spl45_150 | ~ spl45_155)), inference(backward_demodulation, [], [f1389, f1411])).
fof(f1411, plain, ((e0 = op(e3, op(e3, e3))) | ~ spl45_155), inference(avatar_component_clause, [], [f1410])).
fof(f1389, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | spl45_150), inference(avatar_component_clause, [], [f1387])).
fof(f4715, plain, (~ spl45_3 | ~ spl45_6 | ~ spl45_33 | spl45_154), inference(avatar_contradiction_clause, [], [f4714])).
fof(f4714, plain, ($false | (~ spl45_3 | ~ spl45_6 | ~ spl45_33 | spl45_154)), inference(subsumption_resolution, [], [f4710, f655])).
fof(f4710, plain, (~ (e0 = op(e1, e3)) | (~ spl45_3 | ~ spl45_6 | spl45_154)), inference(backward_demodulation, [], [f4619, f540])).
fof(f540, plain, ((e1 = op(e3, e2)) | ~ spl45_6), inference(avatar_component_clause, [], [f538])).
fof(f4619, plain, (~ (e0 = op(op(e3, e2), e3)) | (~ spl45_3 | spl45_154)), inference(forward_demodulation, [], [f1407, f527])).
fof(f527, plain, ((e2 = op(e3, e3)) | ~ spl45_3), inference(avatar_component_clause, [], [f525])).
fof(f1407, plain, (~ (e0 = op(op(e3, op(e3, e3)), e3)) | spl45_154), inference(avatar_component_clause, [], [f1405])).
fof(f1405, plain, (spl45_154 <=> (e0 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_154])])).
fof(f4698, plain, (~ spl45_5 | ~ spl45_21), inference(avatar_split_clause, [], [f4690, f602, f534])).
fof(f4690, plain, (~ (e0 = op(e3, e2)) | ~ spl45_21), inference(backward_demodulation, [], [f211, f604])).
fof(f604, plain, ((e0 = op(e2, e2)) | ~ spl45_21), inference(avatar_component_clause, [], [f602])).
fof(f4683, plain, (~ spl45_25 | ~ spl45_33 | ~ spl45_85), inference(avatar_split_clause, [], [f4679, f979, f653, f619])).
fof(f619, plain, (spl45_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_25])])).
fof(f4679, plain, (~ (e0 = op(e2, e1)) | (~ spl45_33 | ~ spl45_85)), inference(backward_demodulation, [], [f4420, f655])).
fof(f4420, plain, (~ (op(e1, e3) = op(e2, e1)) | ~ spl45_85), inference(backward_demodulation, [], [f205, f981])).
fof(f4682, plain, (spl45_9 | ~ spl45_33 | ~ spl45_85), inference(avatar_split_clause, [], [f4677, f979, f653, f551])).
fof(f4677, plain, ((e0 = op(e3, e1)) | (~ spl45_33 | ~ spl45_85)), inference(backward_demodulation, [], [f981, f655])).
fof(f4673, plain, (~ spl45_34 | ~ spl45_46), inference(avatar_split_clause, [], [f4669, f708, f657])).
fof(f4669, plain, (~ (e1 = op(e1, e3)) | ~ spl45_46), inference(backward_demodulation, [], [f226, f710])).
fof(f4662, plain, (spl45_31 | ~ spl45_55 | ~ spl45_98), inference(avatar_split_clause, [], [f4656, f1041, f746, f644])).
fof(f4656, plain, ((e2 = op(e2, e0)) | (~ spl45_55 | ~ spl45_98)), inference(backward_demodulation, [], [f1043, f748])).
fof(f4661, plain, (~ spl45_23 | ~ spl45_55), inference(avatar_split_clause, [], [f4652, f746, f610])).
fof(f4652, plain, (~ (e2 = op(e2, e2)) | ~ spl45_55), inference(backward_demodulation, [], [f207, f748])).
fof(f207, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4641, plain, (~ spl45_45 | ~ spl45_61), inference(avatar_split_clause, [], [f4628, f772, f704])).
fof(f4628, plain, (~ (e0 = op(e1, e0)) | ~ spl45_61), inference(backward_demodulation, [], [f194, f774])).
fof(f4626, plain, (~ spl45_49 | ~ spl45_16 | spl45_87), inference(avatar_split_clause, [], [f4625, f988, f580, f721])).
fof(f4625, plain, (~ (e0 = op(e0, e3)) | (~ spl45_16 | spl45_87)), inference(forward_demodulation, [], [f989, f582])).
fof(f989, plain, (~ (e0 = op(e0, op(e3, e0))) | spl45_87), inference(avatar_component_clause, [], [f988])).
fof(f4618, plain, (~ spl45_5 | ~ spl45_3 | spl45_155), inference(avatar_split_clause, [], [f4617, f1410, f525, f534])).
fof(f4617, plain, (~ (e0 = op(e3, e2)) | (~ spl45_3 | spl45_155)), inference(forward_demodulation, [], [f1412, f527])).
fof(f1412, plain, (~ (e0 = op(e3, op(e3, e3))) | spl45_155), inference(avatar_component_clause, [], [f1410])).
fof(f4611, plain, (~ spl45_51 | ~ spl45_3), inference(avatar_split_clause, [], [f4610, f525, f729])).
fof(f4610, plain, (~ (e2 = op(e0, e3)) | ~ spl45_3), inference(forward_demodulation, [], [f214, f527])).
fof(f4601, plain, (~ spl45_9 | ~ spl45_139 | spl45_142), inference(avatar_split_clause, [], [f4412, f1345, f1331, f551])).
fof(f1345, plain, (spl45_142 <=> (e0 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_142])])).
fof(f4412, plain, (~ (e0 = op(e3, e1)) | (~ spl45_139 | spl45_142)), inference(backward_demodulation, [], [f1347, f1332])).
fof(f1332, plain, ((e3 = op(e1, op(e1, e1))) | ~ spl45_139), inference(avatar_component_clause, [], [f1331])).
fof(f1347, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | spl45_142), inference(avatar_component_clause, [], [f1345])).
fof(f4597, plain, (~ spl45_1 | ~ spl45_5), inference(avatar_split_clause, [], [f4593, f534, f517])).
fof(f517, plain, (spl45_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_1])])).
fof(f4593, plain, (~ (e0 = op(e3, e3)) | ~ spl45_5), inference(backward_demodulation, [], [f241, f536])).
fof(f4588, plain, (~ spl45_7 | ~ spl45_23), inference(avatar_split_clause, [], [f4582, f610, f542])).
fof(f4582, plain, (~ (e2 = op(e3, e2)) | ~ spl45_23), inference(backward_demodulation, [], [f211, f612])).
fof(f612, plain, ((e2 = op(e2, e2)) | ~ spl45_23), inference(avatar_component_clause, [], [f610])).
fof(f4579, plain, (~ spl45_1 | ~ spl45_49), inference(avatar_split_clause, [], [f4576, f721, f517])).
fof(f4576, plain, (~ (e0 = op(e3, e3)) | ~ spl45_49), inference(backward_demodulation, [], [f214, f723])).
fof(f4573, plain, (~ spl45_54 | spl45_57 | ~ spl45_97 | ~ spl45_98), inference(avatar_contradiction_clause, [], [f4572])).
fof(f4572, plain, ($false | (~ spl45_54 | spl45_57 | ~ spl45_97 | ~ spl45_98)), inference(subsumption_resolution, [], [f4568, f756])).
fof(f756, plain, (~ (e0 = op(e0, e1)) | spl45_57), inference(avatar_component_clause, [], [f755])).
fof(f4568, plain, ((e0 = op(e0, e1)) | (~ spl45_54 | ~ spl45_97 | ~ spl45_98)), inference(backward_demodulation, [], [f4549, f744])).
fof(f4549, plain, ((e0 = op(e0, op(e0, e2))) | (~ spl45_97 | ~ spl45_98)), inference(forward_demodulation, [], [f1038, f1043])).
fof(f4570, plain, (spl45_30 | ~ spl45_54 | ~ spl45_98), inference(avatar_split_clause, [], [f4566, f1041, f742, f640])).
fof(f4566, plain, ((e1 = op(e2, e0)) | (~ spl45_54 | ~ spl45_98)), inference(backward_demodulation, [], [f1043, f744])).
fof(f4569, plain, (~ spl45_22 | ~ spl45_54), inference(avatar_split_clause, [], [f4561, f742, f606])).
fof(f4561, plain, (~ (e1 = op(e2, e2)) | ~ spl45_54), inference(backward_demodulation, [], [f207, f744])).
fof(f4559, plain, (~ spl45_51 | ~ spl45_63), inference(avatar_split_clause, [], [f4555, f780, f729])).
fof(f4555, plain, (~ (e2 = op(e0, e3)) | ~ spl45_63), inference(backward_demodulation, [], [f220, f782])).
fof(f4531, plain, (~ spl45_56 | ~ spl45_60), inference(avatar_split_clause, [], [f4530, f767, f750])).
fof(f4530, plain, (~ (e3 = op(e0, e2)) | ~ spl45_60), inference(forward_demodulation, [], [f221, f769])).
fof(f4529, plain, (~ spl45_52 | ~ spl45_20), inference(avatar_split_clause, [], [f4528, f597, f733])).
fof(f4528, plain, (~ (e3 = op(e0, e3)) | ~ spl45_20), inference(forward_demodulation, [], [f213, f599])).
fof(f599, plain, ((e3 = op(e2, e3)) | ~ spl45_20), inference(avatar_component_clause, [], [f597])).
fof(f4527, plain, (~ spl45_52 | ~ spl45_16 | spl45_88), inference(avatar_split_clause, [], [f4526, f993, f580, f733])).
fof(f4526, plain, (~ (e3 = op(e0, e3)) | (~ spl45_16 | spl45_88)), inference(forward_demodulation, [], [f994, f582])).
fof(f4514, plain, (~ spl45_1 | spl45_15 | ~ spl45_151), inference(avatar_contradiction_clause, [], [f4513])).
fof(f4513, plain, ($false | (~ spl45_1 | spl45_15 | ~ spl45_151)), inference(subsumption_resolution, [], [f4510, f577])).
fof(f577, plain, (~ (e2 = op(e3, e0)) | spl45_15), inference(avatar_component_clause, [], [f576])).
fof(f4510, plain, ((e2 = op(e3, e0)) | (~ spl45_1 | ~ spl45_151)), inference(backward_demodulation, [], [f1392, f519])).
fof(f519, plain, ((e0 = op(e3, e3)) | ~ spl45_1), inference(avatar_component_clause, [], [f517])).
fof(f1392, plain, ((e2 = op(e3, op(e3, e3))) | ~ spl45_151), inference(avatar_component_clause, [], [f1391])).
fof(f1391, plain, (spl45_151 <=> (e2 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl45_151])])).
fof(f4502, plain, (~ spl45_8 | ~ spl45_16), inference(avatar_split_clause, [], [f4499, f580, f546])).
fof(f4499, plain, (~ (e3 = op(e3, e2)) | ~ spl45_16), inference(backward_demodulation, [], [f237, f582])).
fof(f4498, plain, (~ spl45_8 | ~ spl45_20 | spl45_82), inference(avatar_split_clause, [], [f4496, f965, f597, f546])).
fof(f4496, plain, (~ (e3 = op(e3, e2)) | (~ spl45_20 | spl45_82)), inference(backward_demodulation, [], [f966, f599])).
fof(f4493, plain, (~ spl45_17 | ~ spl45_25), inference(avatar_split_clause, [], [f4487, f619, f585])).
fof(f4487, plain, (~ (e0 = op(e2, e3)) | ~ spl45_25), inference(backward_demodulation, [], [f234, f621])).
fof(f621, plain, ((e0 = op(e2, e1)) | ~ spl45_25), inference(avatar_component_clause, [], [f619])).
fof(f4486, plain, (~ spl45_34 | spl45_42 | ~ spl45_84 | ~ spl45_85), inference(avatar_contradiction_clause, [], [f4485])).
fof(f4485, plain, ($false | (~ spl45_34 | spl45_42 | ~ spl45_84 | ~ spl45_85)), inference(subsumption_resolution, [], [f4482, f692])).
fof(f692, plain, (~ (e1 = op(e1, e1)) | spl45_42), inference(avatar_component_clause, [], [f691])).
fof(f4482, plain, ((e1 = op(e1, e1)) | (~ spl45_34 | ~ spl45_84 | ~ spl45_85)), inference(backward_demodulation, [], [f4424, f659])).
fof(f659, plain, ((e1 = op(e1, e3)) | ~ spl45_34), inference(avatar_component_clause, [], [f657])).
fof(f4424, plain, ((e1 = op(e1, op(e1, e3))) | (~ spl45_84 | ~ spl45_85)), inference(forward_demodulation, [], [f976, f981])).
fof(f976, plain, ((e1 = op(e1, op(e3, e1))) | ~ spl45_84), inference(avatar_component_clause, [], [f974])).
fof(f4484, plain, (spl45_10 | ~ spl45_34 | ~ spl45_85), inference(avatar_split_clause, [], [f4477, f979, f657, f555])).
fof(f4477, plain, ((e1 = op(e3, e1)) | (~ spl45_34 | ~ spl45_85)), inference(backward_demodulation, [], [f981, f659])).
fof(f4483, plain, (~ spl45_2 | ~ spl45_34), inference(avatar_split_clause, [], [f4474, f657, f521])).
fof(f4474, plain, (~ (e1 = op(e3, e3)) | ~ spl45_34), inference(backward_demodulation, [], [f216, f659])).
fof(f4471, plain, (~ spl45_36 | ~ spl45_40), inference(avatar_split_clause, [], [f4467, f682, f665])).
fof(f4467, plain, (~ (e3 = op(e1, e3)) | ~ spl45_40), inference(backward_demodulation, [], [f229, f684])).
fof(f4470, plain, (~ spl45_8 | ~ spl45_40), inference(avatar_split_clause, [], [f4466, f682, f546])).
fof(f4466, plain, (~ (e3 = op(e3, e2)) | ~ spl45_40), inference(backward_demodulation, [], [f210, f684])).
fof(f4465, plain, (~ spl45_39 | ~ spl45_43 | spl45_141), inference(avatar_split_clause, [], [f4460, f1340, f695, f678])).
fof(f1340, plain, (spl45_141 <=> (e2 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl45_141])])).
fof(f4460, plain, (~ (e2 = op(e1, e2)) | (~ spl45_43 | spl45_141)), inference(backward_demodulation, [], [f1342, f697])).
fof(f1342, plain, (~ (e2 = op(e1, op(e1, e1))) | spl45_141), inference(avatar_component_clause, [], [f1340])).
fof(f4464, plain, (spl45_40 | ~ spl45_43 | ~ spl45_139), inference(avatar_split_clause, [], [f4459, f1331, f695, f682])).
fof(f4459, plain, ((e3 = op(e1, e2)) | (~ spl45_43 | ~ spl45_139)), inference(backward_demodulation, [], [f1332, f697])).
fof(f4463, plain, (~ spl45_39 | ~ spl45_43), inference(avatar_split_clause, [], [f4456, f695, f678])).
fof(f4456, plain, (~ (e2 = op(e1, e2)) | ~ spl45_43), inference(backward_demodulation, [], [f227, f697])).
fof(f4462, plain, (~ spl45_11 | ~ spl45_43), inference(avatar_split_clause, [], [f4455, f695, f559])).
fof(f4455, plain, (~ (e2 = op(e3, e1)) | ~ spl45_43), inference(backward_demodulation, [], [f204, f697])).
fof(f4453, plain, (~ spl45_41 | ~ spl45_45), inference(avatar_split_clause, [], [f4445, f704, f687])).
fof(f4445, plain, (~ (e0 = op(e1, e1)) | ~ spl45_45), inference(backward_demodulation, [], [f224, f706])).
fof(f4452, plain, (~ spl45_13 | ~ spl45_45), inference(avatar_split_clause, [], [f4444, f704, f568])).
fof(f568, plain, (spl45_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_13])])).
fof(f4444, plain, (~ (e0 = op(e3, e0)) | ~ spl45_45), inference(backward_demodulation, [], [f198, f706])).
fof(f4443, plain, (spl45_34 | ~ spl45_60 | ~ spl45_113), inference(avatar_split_clause, [], [f4441, f1114, f767, f657])).
fof(f4441, plain, ((e1 = op(e1, e3)) | (~ spl45_60 | ~ spl45_113)), inference(backward_demodulation, [], [f1116, f769])).
fof(f4442, plain, (~ spl45_28 | ~ spl45_60), inference(avatar_split_clause, [], [f4438, f767, f631])).
fof(f4438, plain, (~ (e3 = op(e2, e1)) | ~ spl45_60), inference(backward_demodulation, [], [f201, f769])).
fof(f4411, plain, (~ spl45_28 | ~ spl45_22 | spl45_145), inference(avatar_split_clause, [], [f4410, f1361, f606, f631])).
fof(f1361, plain, (spl45_145 <=> (e3 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl45_145])])).
fof(f4410, plain, (~ (e3 = op(e2, e1)) | (~ spl45_22 | spl45_145)), inference(forward_demodulation, [], [f1363, f608])).
fof(f1363, plain, (~ (e3 = op(e2, op(e2, e2))) | spl45_145), inference(avatar_component_clause, [], [f1361])).
fof(f4385, plain, (~ spl45_9 | ~ spl45_13), inference(avatar_split_clause, [], [f4379, f568, f551])).
fof(f4379, plain, (~ (e0 = op(e3, e1)) | ~ spl45_13), inference(backward_demodulation, [], [f236, f570])).
fof(f570, plain, ((e0 = op(e3, e0)) | ~ spl45_13), inference(avatar_component_clause, [], [f568])).
fof(f4378, plain, (~ spl45_9 | ~ spl45_17 | ~ spl45_82), inference(avatar_split_clause, [], [f4373, f965, f585, f551])).
fof(f4373, plain, (~ (e0 = op(e3, e1)) | (~ spl45_17 | ~ spl45_82)), inference(backward_demodulation, [], [f4277, f587])).
fof(f4277, plain, (~ (op(e2, e3) = op(e3, e1)) | ~ spl45_82), inference(forward_demodulation, [], [f239, f967])).
fof(f4377, plain, (~ spl45_13 | ~ spl45_17 | ~ spl45_82), inference(avatar_split_clause, [], [f4372, f965, f585, f568])).
fof(f4372, plain, (~ (e0 = op(e3, e0)) | (~ spl45_17 | ~ spl45_82)), inference(backward_demodulation, [], [f4276, f587])).
fof(f4276, plain, (~ (op(e2, e3) = op(e3, e0)) | ~ spl45_82), inference(forward_demodulation, [], [f237, f967])).
fof(f4376, plain, (spl45_5 | ~ spl45_17 | ~ spl45_82), inference(avatar_contradiction_clause, [], [f4375])).
fof(f4375, plain, ($false | (spl45_5 | ~ spl45_17 | ~ spl45_82)), inference(subsumption_resolution, [], [f4369, f535])).
fof(f535, plain, (~ (e0 = op(e3, e2)) | spl45_5), inference(avatar_component_clause, [], [f534])).
fof(f4369, plain, ((e0 = op(e3, e2)) | (~ spl45_17 | ~ spl45_82)), inference(backward_demodulation, [], [f967, f587])).
fof(f4365, plain, (~ spl45_18 | ~ spl45_22), inference(avatar_split_clause, [], [f4358, f606, f589])).
fof(f4358, plain, (~ (e1 = op(e2, e3)) | ~ spl45_22), inference(backward_demodulation, [], [f235, f608])).
fof(f235, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4356, plain, (~ spl45_23 | ~ spl45_31), inference(avatar_split_clause, [], [f4354, f644, f610])).
fof(f4354, plain, (~ (e2 = op(e2, e2)) | ~ spl45_31), inference(backward_demodulation, [], [f231, f646])).
fof(f646, plain, ((e2 = op(e2, e0)) | ~ spl45_31), inference(avatar_component_clause, [], [f644])).
fof(f4349, plain, (~ spl45_23 | ~ spl45_39), inference(avatar_split_clause, [], [f4345, f678, f610])).
fof(f4345, plain, (~ (e2 = op(e2, e2)) | ~ spl45_39), inference(backward_demodulation, [], [f209, f680])).
fof(f4328, plain, (~ spl45_31 | ~ spl45_53 | spl45_111), inference(avatar_split_clause, [], [f4325, f1104, f738, f644])).
fof(f4325, plain, (~ (e2 = op(e2, e0)) | (~ spl45_53 | spl45_111)), inference(backward_demodulation, [], [f1105, f740])).
fof(f1105, plain, (~ (e2 = op(e2, op(e0, e2))) | spl45_111), inference(avatar_component_clause, [], [f1104])).
fof(f4315, plain, (spl45_8 | ~ spl45_51 | ~ spl45_109), inference(avatar_split_clause, [], [f4314, f1094, f729, f546])).
fof(f4314, plain, ((e3 = op(e3, e2)) | (~ spl45_51 | ~ spl45_109)), inference(forward_demodulation, [], [f1096, f731])).
fof(f4313, plain, (~ spl45_6 | spl45_144 | ~ spl45_145), inference(avatar_split_clause, [], [f4310, f1361, f1357, f538])).
fof(f1357, plain, (spl45_144 <=> (e1 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_144])])).
fof(f4310, plain, (~ (e1 = op(e3, e2)) | (spl45_144 | ~ spl45_145)), inference(backward_demodulation, [], [f1359, f1362])).
fof(f1362, plain, ((e3 = op(e2, op(e2, e2))) | ~ spl45_145), inference(avatar_component_clause, [], [f1361])).
fof(f1359, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | spl45_144), inference(avatar_component_clause, [], [f1357])).
fof(f4308, plain, (~ spl45_17 | ~ spl45_151 | spl45_154), inference(avatar_split_clause, [], [f4306, f1405, f1391, f585])).
fof(f4306, plain, (~ (e0 = op(e2, e3)) | (~ spl45_151 | spl45_154)), inference(backward_demodulation, [], [f1407, f1392])).
fof(f4280, plain, (~ spl45_24 | ~ spl45_28), inference(avatar_split_clause, [], [f4279, f631, f614])).
fof(f4279, plain, (~ (e3 = op(e2, e2)) | ~ spl45_28), inference(forward_demodulation, [], [f233, f633])).
fof(f4273, plain, (~ spl45_3 | ~ spl45_4 | spl45_151), inference(avatar_split_clause, [], [f4268, f1391, f529, f525])).
fof(f4268, plain, (~ (e2 = op(e3, e3)) | (~ spl45_4 | spl45_151)), inference(backward_demodulation, [], [f1393, f531])).
fof(f531, plain, ((e3 = op(e3, e3)) | ~ spl45_4), inference(avatar_component_clause, [], [f529])).
fof(f1393, plain, (~ (e2 = op(e3, op(e3, e3))) | spl45_151), inference(avatar_component_clause, [], [f1391])).
fof(f4261, plain, (~ spl45_3 | ~ spl45_15), inference(avatar_split_clause, [], [f4259, f576, f525])).
fof(f4259, plain, (~ (e2 = op(e3, e3)) | ~ spl45_15), inference(backward_demodulation, [], [f238, f578])).
fof(f4252, plain, (~ spl45_12 | ~ spl45_28), inference(avatar_split_clause, [], [f4250, f631, f563])).
fof(f4250, plain, (~ (e3 = op(e3, e1)) | ~ spl45_28), inference(backward_demodulation, [], [f205, f633])).
fof(f4247, plain, (~ spl45_25 | ~ spl45_29), inference(avatar_split_clause, [], [f4241, f636, f619])).
fof(f4241, plain, (~ (e0 = op(e2, e1)) | ~ spl45_29), inference(backward_demodulation, [], [f230, f638])).
fof(f4246, plain, (~ spl45_13 | ~ spl45_29), inference(avatar_split_clause, [], [f4240, f636, f568])).
fof(f4240, plain, (~ (e0 = op(e3, e0)) | ~ spl45_29), inference(backward_demodulation, [], [f199, f638])).
fof(f4232, plain, (~ spl45_30 | ~ spl45_46), inference(avatar_split_clause, [], [f4226, f708, f640])).
fof(f4226, plain, (~ (e1 = op(e2, e0)) | ~ spl45_46), inference(backward_demodulation, [], [f197, f710])).
fof(f4225, plain, (~ spl45_49 | ~ spl45_51), inference(avatar_contradiction_clause, [], [f4224])).
fof(f4224, plain, ($false | (~ spl45_49 | ~ spl45_51)), inference(subsumption_resolution, [], [f4223, f243])).
fof(f243, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f4223, plain, ((e0 = e2) | (~ spl45_49 | ~ spl45_51)), inference(forward_demodulation, [], [f731, f723])).
fof(f4212, plain, (~ spl45_6 | spl45_27 | ~ spl45_81), inference(avatar_contradiction_clause, [], [f4211])).
fof(f4211, plain, ($false | (~ spl45_6 | spl45_27 | ~ spl45_81)), inference(subsumption_resolution, [], [f4210, f628])).
fof(f628, plain, (~ (e2 = op(e2, e1)) | spl45_27), inference(avatar_component_clause, [], [f627])).
fof(f4210, plain, ((e2 = op(e2, e1)) | (~ spl45_6 | ~ spl45_81)), inference(forward_demodulation, [], [f962, f540])).
fof(f4209, plain, (spl45_18 | ~ spl45_6 | ~ spl45_82), inference(avatar_split_clause, [], [f4208, f965, f538, f589])).
fof(f4208, plain, ((e1 = op(e2, e3)) | (~ spl45_6 | ~ spl45_82)), inference(forward_demodulation, [], [f967, f540])).
fof(f4207, plain, (~ spl45_25 | ~ spl45_37 | spl45_95), inference(avatar_split_clause, [], [f4206, f1027, f670, f619])).
fof(f4206, plain, (~ (e0 = op(e2, e1)) | (~ spl45_37 | spl45_95)), inference(forward_demodulation, [], [f1028, f672])).
fof(f4205, plain, (spl45_46 | ~ spl45_58 | ~ spl45_107), inference(avatar_split_clause, [], [f4204, f1085, f759, f708])).
fof(f4204, plain, ((e1 = op(e1, e0)) | (~ spl45_58 | ~ spl45_107)), inference(forward_demodulation, [], [f1087, f761])).
fof(f4202, plain, (~ spl45_30 | ~ spl45_135 | spl45_136), inference(avatar_split_clause, [], [f4200, f1315, f1310, f640])).
fof(f1315, plain, (spl45_136 <=> (e1 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_136])])).
fof(f4200, plain, (~ (e1 = op(e2, e0)) | (~ spl45_135 | spl45_136)), inference(backward_demodulation, [], [f1317, f1311])).
fof(f1317, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | spl45_136), inference(avatar_component_clause, [], [f1315])).
fof(f4192, plain, (spl45_13 | ~ spl45_49 | ~ spl45_88), inference(avatar_split_clause, [], [f4191, f993, f721, f568])).
fof(f4191, plain, ((e0 = op(e3, e0)) | (~ spl45_49 | ~ spl45_88)), inference(forward_demodulation, [], [f995, f723])).
fof(f4190, plain, (~ spl45_16 | ~ spl45_49 | spl45_109), inference(avatar_split_clause, [], [f4138, f1094, f721, f580])).
fof(f4138, plain, (~ (e3 = op(e3, e0)) | (~ spl45_49 | spl45_109)), inference(backward_demodulation, [], [f1095, f723])).
fof(f1095, plain, (~ (e3 = op(e3, op(e0, e3))) | spl45_109), inference(avatar_component_clause, [], [f1094])).
fof(f4185, plain, (~ spl45_13 | spl45_61 | ~ spl45_87), inference(avatar_contradiction_clause, [], [f4184])).
fof(f4184, plain, ($false | (~ spl45_13 | spl45_61 | ~ spl45_87)), inference(subsumption_resolution, [], [f4183, f773])).
fof(f4183, plain, ((e0 = op(e0, e0)) | (~ spl45_13 | ~ spl45_87)), inference(forward_demodulation, [], [f990, f570])).
fof(f4176, plain, (~ spl45_48 | ~ spl45_58 | ~ spl45_107), inference(avatar_contradiction_clause, [], [f4175])).
fof(f4175, plain, ($false | (~ spl45_48 | ~ spl45_58 | ~ spl45_107)), inference(subsumption_resolution, [], [f4174, f246])).
fof(f246, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f4174, plain, ((e1 = e3) | (~ spl45_48 | ~ spl45_58 | ~ spl45_107)), inference(forward_demodulation, [], [f4173, f761])).
fof(f4173, plain, ((e3 = op(e0, e1)) | (~ spl45_48 | ~ spl45_107)), inference(forward_demodulation, [], [f1087, f718])).
fof(f718, plain, ((e3 = op(e1, e0)) | ~ spl45_48), inference(avatar_component_clause, [], [f716])).
fof(f4159, plain, (~ spl45_26 | ~ spl45_30), inference(avatar_split_clause, [], [f4155, f640, f623])).
fof(f4155, plain, (~ (e1 = op(e2, e1)) | ~ spl45_30), inference(backward_demodulation, [], [f230, f642])).
fof(f4148, plain, (~ spl45_36 | ~ spl45_48), inference(avatar_split_clause, [], [f4146, f716, f665])).
fof(f4146, plain, (~ (e3 = op(e1, e3)) | ~ spl45_48), inference(backward_demodulation, [], [f226, f718])).
fof(f4147, plain, (~ spl45_32 | ~ spl45_48), inference(avatar_split_clause, [], [f4145, f716, f648])).
fof(f4145, plain, (~ (e3 = op(e2, e0)) | ~ spl45_48), inference(backward_demodulation, [], [f197, f718])).
fof(f4136, plain, (spl45_42 | ~ spl45_58 | ~ spl45_113), inference(avatar_contradiction_clause, [], [f4135])).
fof(f4135, plain, ($false | (spl45_42 | ~ spl45_58 | ~ spl45_113)), inference(subsumption_resolution, [], [f4131, f692])).
fof(f4131, plain, ((e1 = op(e1, e1)) | (~ spl45_58 | ~ spl45_113)), inference(backward_demodulation, [], [f1116, f761])).
fof(f4134, plain, (~ spl45_46 | ~ spl45_58 | spl45_107), inference(avatar_split_clause, [], [f4130, f1085, f759, f708])).
fof(f4130, plain, (~ (e1 = op(e1, e0)) | (~ spl45_58 | spl45_107)), inference(backward_demodulation, [], [f1086, f761])).
fof(f4123, plain, (spl45_49 | ~ spl45_13 | ~ spl45_88), inference(avatar_split_clause, [], [f4122, f993, f568, f721])).
fof(f4122, plain, ((e0 = op(e0, e3)) | (~ spl45_13 | ~ spl45_88)), inference(forward_demodulation, [], [f995, f570])).
fof(f4121, plain, (spl45_25 | ~ spl45_37 | ~ spl45_95), inference(avatar_split_clause, [], [f4120, f1027, f670, f619])).
fof(f4120, plain, ((e0 = op(e2, e1)) | (~ spl45_37 | ~ spl45_95)), inference(forward_demodulation, [], [f1029, f672])).
fof(f4117, plain, (spl45_31 | ~ spl45_37 | ~ spl45_102), inference(avatar_contradiction_clause, [], [f4116])).
fof(f4116, plain, ($false | (spl45_31 | ~ spl45_37 | ~ spl45_102)), inference(subsumption_resolution, [], [f4115, f645])).
fof(f4115, plain, ((e2 = op(e2, e0)) | (~ spl45_37 | ~ spl45_102)), inference(forward_demodulation, [], [f1062, f672])).
fof(f1062, plain, ((e2 = op(e2, op(e1, e2))) | ~ spl45_102), inference(avatar_component_clause, [], [f1060])).
fof(f4107, plain, (~ spl45_59 | ~ spl45_63), inference(avatar_split_clause, [], [f4106, f780, f763])).
fof(f4106, plain, (~ (e2 = op(e0, e1)) | ~ spl45_63), inference(forward_demodulation, [], [f218, f782])).
fof(f4104, plain, (~ spl45_60 | ~ spl45_37 | ~ spl45_43 | spl45_140), inference(avatar_split_clause, [], [f4028, f1336, f695, f670, f767])).
fof(f1336, plain, (spl45_140 <=> (e3 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl45_140])])).
fof(f4028, plain, (~ (e3 = op(e0, e1)) | (~ spl45_37 | ~ spl45_43 | spl45_140)), inference(forward_demodulation, [], [f4027, f672])).
fof(f4027, plain, (~ (e3 = op(op(e1, e2), e1)) | (~ spl45_43 | spl45_140)), inference(forward_demodulation, [], [f1338, f697])).
fof(f1338, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | spl45_140), inference(avatar_component_clause, [], [f1336])).
fof(f4103, plain, (~ spl45_52 | ~ spl45_56), inference(avatar_split_clause, [], [f4102, f750, f733])).
fof(f4102, plain, (~ (e3 = op(e0, e3)) | ~ spl45_56), inference(forward_demodulation, [], [f223, f752])).
fof(f4097, plain, (~ spl45_47 | ~ spl45_63), inference(avatar_split_clause, [], [f4096, f780, f712])).
fof(f4096, plain, (~ (e2 = op(e1, e0)) | ~ spl45_63), inference(forward_demodulation, [], [f194, f782])).
fof(f4090, plain, (~ spl45_19 | ~ spl45_23), inference(avatar_split_clause, [], [f4089, f610, f593])).
fof(f4089, plain, (~ (e2 = op(e2, e3)) | ~ spl45_23), inference(forward_demodulation, [], [f235, f612])).
fof(f4087, plain, (~ spl45_32 | spl45_49 | ~ spl45_97), inference(avatar_contradiction_clause, [], [f4086])).
fof(f4086, plain, ($false | (~ spl45_32 | spl45_49 | ~ spl45_97)), inference(subsumption_resolution, [], [f4085, f722])).
fof(f722, plain, (~ (e0 = op(e0, e3)) | spl45_49), inference(avatar_component_clause, [], [f721])).
fof(f4085, plain, ((e0 = op(e0, e3)) | (~ spl45_32 | ~ spl45_97)), inference(forward_demodulation, [], [f1038, f650])).
fof(f4051, plain, (spl45_32 | ~ spl45_56 | ~ spl45_98), inference(avatar_split_clause, [], [f4046, f1041, f750, f648])).
fof(f4046, plain, ((e3 = op(e2, e0)) | (~ spl45_56 | ~ spl45_98)), inference(backward_demodulation, [], [f1043, f752])).
fof(f4050, plain, (~ spl45_24 | ~ spl45_56), inference(avatar_split_clause, [], [f4045, f750, f614])).
fof(f4045, plain, (~ (e3 = op(e2, e2)) | ~ spl45_56), inference(backward_demodulation, [], [f207, f752])).
fof(f4043, plain, (~ spl45_56 | ~ spl45_63 | spl45_133), inference(avatar_split_clause, [], [f4037, f1301, f780, f750])).
fof(f1301, plain, (spl45_133 <=> (e3 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl45_133])])).
fof(f4037, plain, (~ (e3 = op(e0, e2)) | (~ spl45_63 | spl45_133)), inference(backward_demodulation, [], [f1303, f782])).
fof(f1303, plain, (~ (e3 = op(e0, op(e0, e0))) | spl45_133), inference(avatar_component_clause, [], [f1301])).
fof(f4005, plain, (~ spl45_56 | spl45_146 | ~ spl45_149), inference(avatar_split_clause, [], [f3877, f1380, f1366, f750])).
fof(f1366, plain, (spl45_146 <=> (e3 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl45_146])])).
fof(f3877, plain, (~ (e3 = op(e0, e2)) | (spl45_146 | ~ spl45_149)), inference(forward_demodulation, [], [f1368, f1381])).
fof(f1381, plain, ((e0 = op(e2, op(e2, e2))) | ~ spl45_149), inference(avatar_component_clause, [], [f1380])).
fof(f1368, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | spl45_146), inference(avatar_component_clause, [], [f1366])).
fof(f3993, plain, (~ spl45_5 | ~ spl45_13), inference(avatar_split_clause, [], [f3988, f568, f534])).
fof(f3988, plain, (~ (e0 = op(e3, e2)) | ~ spl45_13), inference(backward_demodulation, [], [f237, f570])).
fof(f3981, plain, (~ spl45_20 | ~ spl45_24), inference(avatar_split_clause, [], [f3979, f614, f597])).
fof(f3979, plain, (~ (e3 = op(e2, e3)) | ~ spl45_24), inference(backward_demodulation, [], [f235, f616])).
fof(f3958, plain, (~ spl45_5 | ~ spl45_37), inference(avatar_split_clause, [], [f3950, f670, f534])).
fof(f3950, plain, (~ (e0 = op(e3, e2)) | ~ spl45_37), inference(backward_demodulation, [], [f210, f672])).
fof(f3928, plain, (~ spl45_49 | ~ spl45_57), inference(avatar_split_clause, [], [f3921, f755, f721])).
fof(f3921, plain, (~ (e0 = op(e0, e3)) | ~ spl45_57), inference(backward_demodulation, [], [f222, f757])).
fof(f3927, plain, (~ spl45_25 | ~ spl45_57), inference(avatar_split_clause, [], [f3920, f755, f619])).
fof(f3920, plain, (~ (e0 = op(e2, e1)) | ~ spl45_57), inference(backward_demodulation, [], [f201, f757])).
fof(f3919, plain, (spl45_23 | ~ spl45_55 | ~ spl45_111), inference(avatar_contradiction_clause, [], [f3918])).
fof(f3918, plain, ($false | (spl45_23 | ~ spl45_55 | ~ spl45_111)), inference(subsumption_resolution, [], [f3917, f611])).
fof(f3917, plain, ((e2 = op(e2, e2)) | (~ spl45_55 | ~ spl45_111)), inference(forward_demodulation, [], [f1106, f748])).
fof(f3906, plain, (~ spl45_3 | spl45_7 | ~ spl45_151), inference(avatar_contradiction_clause, [], [f3905])).
fof(f3905, plain, ($false | (~ spl45_3 | spl45_7 | ~ spl45_151)), inference(subsumption_resolution, [], [f3904, f543])).
fof(f543, plain, (~ (e2 = op(e3, e2)) | spl45_7), inference(avatar_component_clause, [], [f542])).
fof(f3904, plain, ((e2 = op(e3, e2)) | (~ spl45_3 | ~ spl45_151)), inference(forward_demodulation, [], [f1392, f527])).
fof(f3899, plain, (~ spl45_52 | ~ spl45_64), inference(avatar_split_clause, [], [f3898, f784, f733])).
fof(f3898, plain, (~ (e3 = op(e0, e3)) | ~ spl45_64), inference(forward_demodulation, [], [f220, f786])).
fof(f3890, plain, (~ spl45_37 | ~ spl45_43 | spl45_143), inference(avatar_split_clause, [], [f3846, f1350, f695, f670])).
fof(f3846, plain, (~ (e0 = op(e1, e2)) | (~ spl45_43 | spl45_143)), inference(backward_demodulation, [], [f1352, f697])).
fof(f3888, plain, (spl45_36 | ~ spl45_12 | ~ spl45_85), inference(avatar_split_clause, [], [f3887, f979, f563, f665])).
fof(f3887, plain, ((e3 = op(e1, e3)) | (~ spl45_12 | ~ spl45_85)), inference(forward_demodulation, [], [f981, f565])).
fof(f3884, plain, (~ spl45_27 | ~ spl45_31), inference(avatar_split_clause, [], [f3883, f644, f627])).
fof(f3883, plain, (~ (e2 = op(e2, e1)) | ~ spl45_31), inference(forward_demodulation, [], [f230, f646])).
fof(f3871, plain, (~ spl45_7 | ~ spl45_3), inference(avatar_split_clause, [], [f3870, f525, f542])).
fof(f3870, plain, (~ (e2 = op(e3, e2)) | ~ spl45_3), inference(forward_demodulation, [], [f241, f527])).
fof(f3869, plain, (~ spl45_8 | ~ spl45_12), inference(avatar_split_clause, [], [f3868, f563, f546])).
fof(f3868, plain, (~ (e3 = op(e3, e2)) | ~ spl45_12), inference(forward_demodulation, [], [f239, f565])).
fof(f3860, plain, (~ spl45_14 | ~ spl45_34 | ~ spl45_85), inference(avatar_split_clause, [], [f3855, f979, f657, f572])).
fof(f3855, plain, (~ (e1 = op(e3, e0)) | (~ spl45_34 | ~ spl45_85)), inference(backward_demodulation, [], [f3797, f659])).
fof(f3797, plain, (~ (op(e1, e3) = op(e3, e0)) | ~ spl45_85), inference(forward_demodulation, [], [f236, f981])).
fof(f3859, plain, (spl45_12 | ~ spl45_34 | ~ spl45_100), inference(avatar_split_clause, [], [f3853, f1050, f657, f563])).
fof(f3853, plain, ((e3 = op(e3, e1)) | (~ spl45_34 | ~ spl45_100)), inference(backward_demodulation, [], [f1052, f659])).
fof(f3784, plain, (~ spl45_31 | spl45_53 | ~ spl45_97), inference(avatar_contradiction_clause, [], [f3783])).
fof(f3783, plain, ($false | (~ spl45_31 | spl45_53 | ~ spl45_97)), inference(subsumption_resolution, [], [f3782, f739])).
fof(f3782, plain, ((e0 = op(e0, e2)) | (~ spl45_31 | ~ spl45_97)), inference(forward_demodulation, [], [f1038, f646])).
fof(f3769, plain, (~ spl45_22 | ~ spl45_25 | spl45_149), inference(avatar_contradiction_clause, [], [f3768])).
fof(f3768, plain, ($false | (~ spl45_22 | ~ spl45_25 | spl45_149)), inference(subsumption_resolution, [], [f3761, f621])).
fof(f3761, plain, (~ (e0 = op(e2, e1)) | (~ spl45_22 | spl45_149)), inference(backward_demodulation, [], [f1382, f608])).
fof(f3751, plain, (spl45_8 | ~ spl45_35 | ~ spl45_100), inference(avatar_split_clause, [], [f3750, f1050, f661, f546])).
fof(f3750, plain, ((e3 = op(e3, e2)) | (~ spl45_35 | ~ spl45_100)), inference(backward_demodulation, [], [f1052, f663])).
fof(f3720, plain, (spl45_16 | ~ spl45_49 | ~ spl45_109), inference(avatar_split_clause, [], [f3717, f1094, f721, f580])).
fof(f3717, plain, ((e3 = op(e3, e0)) | (~ spl45_49 | ~ spl45_109)), inference(backward_demodulation, [], [f1096, f723])).
fof(f3719, plain, (~ spl45_13 | ~ spl45_49 | spl45_88), inference(avatar_split_clause, [], [f3716, f993, f721, f568])).
fof(f3716, plain, (~ (e0 = op(e3, e0)) | (~ spl45_49 | spl45_88)), inference(backward_demodulation, [], [f994, f723])).
fof(f3684, plain, (spl45_35 | ~ spl45_11 | ~ spl45_85), inference(avatar_split_clause, [], [f3683, f979, f559, f661])).
fof(f3683, plain, ((e2 = op(e1, e3)) | (~ spl45_11 | ~ spl45_85)), inference(forward_demodulation, [], [f981, f561])).
fof(f3679, plain, (~ spl45_48 | ~ spl45_60 | spl45_107), inference(avatar_split_clause, [], [f3678, f1085, f767, f716])).
fof(f3678, plain, (~ (e3 = op(e1, e0)) | (~ spl45_60 | spl45_107)), inference(forward_demodulation, [], [f1086, f769])).
fof(f3662, plain, (spl45_38 | ~ spl45_11 | ~ spl45_84), inference(avatar_split_clause, [], [f3634, f974, f559, f674])).
fof(f3634, plain, ((e1 = op(e1, e2)) | (~ spl45_11 | ~ spl45_84)), inference(backward_demodulation, [], [f976, f561])).
fof(f3654, plain, (~ spl45_15 | ~ spl45_11), inference(avatar_split_clause, [], [f3653, f559, f576])).
fof(f3653, plain, (~ (e2 = op(e3, e0)) | ~ spl45_11), inference(forward_demodulation, [], [f236, f561])).
fof(f3647, plain, (~ spl45_2 | ~ spl45_11 | spl45_151), inference(avatar_contradiction_clause, [], [f3646])).
fof(f3646, plain, ($false | (~ spl45_2 | ~ spl45_11 | spl45_151)), inference(subsumption_resolution, [], [f3641, f561])).
fof(f3641, plain, (~ (e2 = op(e3, e1)) | (~ spl45_2 | spl45_151)), inference(backward_demodulation, [], [f1393, f523])).
fof(f523, plain, ((e1 = op(e3, e3)) | ~ spl45_2), inference(avatar_component_clause, [], [f521])).
fof(f3637, plain, (~ spl45_3 | ~ spl45_11), inference(avatar_split_clause, [], [f3635, f559, f525])).
fof(f3635, plain, (~ (e2 = op(e3, e3)) | ~ spl45_11), inference(backward_demodulation, [], [f240, f561])).
fof(f3629, plain, (~ spl45_58 | ~ spl45_60), inference(avatar_contradiction_clause, [], [f3628])).
fof(f3628, plain, ($false | (~ spl45_58 | ~ spl45_60)), inference(subsumption_resolution, [], [f3627, f246])).
fof(f3627, plain, ((e1 = e3) | (~ spl45_58 | ~ spl45_60)), inference(backward_demodulation, [], [f769, f761])).
fof(f3604, plain, (~ spl45_30 | ~ spl45_53 | ~ spl45_98), inference(avatar_contradiction_clause, [], [f3603])).
fof(f3603, plain, ($false | (~ spl45_30 | ~ spl45_53 | ~ spl45_98)), inference(subsumption_resolution, [], [f3602, f242])).
fof(f242, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f3602, plain, ((e0 = e1) | (~ spl45_30 | ~ spl45_53 | ~ spl45_98)), inference(forward_demodulation, [], [f3601, f740])).
fof(f3601, plain, ((e1 = op(e0, e2)) | (~ spl45_30 | ~ spl45_98)), inference(forward_demodulation, [], [f1043, f642])).
fof(f3591, plain, (~ spl45_48 | spl45_49 | ~ spl45_106), inference(avatar_contradiction_clause, [], [f3590])).
fof(f3590, plain, ($false | (~ spl45_48 | spl45_49 | ~ spl45_106)), inference(subsumption_resolution, [], [f3589, f722])).
fof(f3589, plain, ((e0 = op(e0, e3)) | (~ spl45_48 | ~ spl45_106)), inference(forward_demodulation, [], [f1082, f718])).
fof(f3568, plain, (~ spl45_38 | ~ spl45_43 | spl45_138), inference(avatar_contradiction_clause, [], [f3567])).
fof(f3567, plain, ($false | (~ spl45_38 | ~ spl45_43 | spl45_138)), inference(subsumption_resolution, [], [f3566, f697])).
fof(f3566, plain, (~ (e2 = op(e1, e1)) | (~ spl45_38 | ~ spl45_43 | spl45_138)), inference(forward_demodulation, [], [f3474, f676])).
fof(f3474, plain, (~ (e2 = op(op(e1, e2), e1)) | (~ spl45_43 | spl45_138)), inference(forward_demodulation, [], [f1329, f697])).
fof(f3554, plain, (spl45_4 | ~ spl45_20 | ~ spl45_90), inference(avatar_split_clause, [], [f3550, f1002, f597, f529])).
fof(f3550, plain, ((e3 = op(e3, e3)) | (~ spl45_20 | ~ spl45_90)), inference(backward_demodulation, [], [f1004, f599])).
fof(f3553, plain, (spl45_8 | ~ spl45_20 | ~ spl45_82), inference(avatar_split_clause, [], [f3549, f965, f597, f546])).
fof(f3549, plain, ((e3 = op(e3, e2)) | (~ spl45_20 | ~ spl45_82)), inference(backward_demodulation, [], [f967, f599])).
fof(f3552, plain, (~ spl45_4 | ~ spl45_20), inference(avatar_split_clause, [], [f3548, f597, f529])).
fof(f3548, plain, (~ (e3 = op(e3, e3)) | ~ spl45_20), inference(backward_demodulation, [], [f217, f599])).
fof(f3545, plain, (~ spl45_23 | ~ spl45_145), inference(avatar_contradiction_clause, [], [f3544])).
fof(f3544, plain, ($false | (~ spl45_23 | ~ spl45_145)), inference(subsumption_resolution, [], [f3543, f247])).
fof(f3543, plain, ((e2 = e3) | (~ spl45_23 | ~ spl45_145)), inference(forward_demodulation, [], [f3538, f612])).
fof(f3538, plain, ((e3 = op(e2, e2)) | (~ spl45_23 | ~ spl45_145)), inference(backward_demodulation, [], [f1362, f612])).
fof(f3534, plain, (~ spl45_21 | ~ spl45_25), inference(avatar_split_clause, [], [f3529, f619, f602])).
fof(f3529, plain, (~ (e0 = op(e2, e2)) | ~ spl45_25), inference(backward_demodulation, [], [f233, f621])).
fof(f3528, plain, (~ spl45_22 | ~ spl45_30), inference(avatar_split_clause, [], [f3524, f640, f606])).
fof(f3524, plain, (~ (e1 = op(e2, e2)) | ~ spl45_30), inference(backward_demodulation, [], [f231, f642])).
fof(f3512, plain, (~ spl45_40 | ~ spl45_48), inference(avatar_split_clause, [], [f3511, f716, f682])).
fof(f3511, plain, (~ (e3 = op(e1, e2)) | ~ spl45_48), inference(backward_demodulation, [], [f225, f718])).
fof(f3509, plain, (~ spl45_21 | ~ spl45_53), inference(avatar_split_clause, [], [f3503, f738, f602])).
fof(f3503, plain, (~ (e0 = op(e2, e2)) | ~ spl45_53), inference(backward_demodulation, [], [f207, f740])).
fof(f3501, plain, (~ spl45_33 | ~ spl45_60 | ~ spl45_113), inference(avatar_contradiction_clause, [], [f3500])).
fof(f3500, plain, ($false | (~ spl45_33 | ~ spl45_60 | ~ spl45_113)), inference(subsumption_resolution, [], [f3499, f242])).
fof(f3499, plain, ((e0 = e1) | (~ spl45_33 | ~ spl45_60 | ~ spl45_113)), inference(forward_demodulation, [], [f3496, f655])).
fof(f3496, plain, ((e1 = op(e1, e3)) | (~ spl45_60 | ~ spl45_113)), inference(backward_demodulation, [], [f1116, f769])).
fof(f3453, plain, (spl45_1 | ~ spl45_4 | ~ spl45_155), inference(avatar_contradiction_clause, [], [f3452])).
fof(f3452, plain, ($false | (spl45_1 | ~ spl45_4 | ~ spl45_155)), inference(subsumption_resolution, [], [f3451, f518])).
fof(f518, plain, (~ (e0 = op(e3, e3)) | spl45_1), inference(avatar_component_clause, [], [f517])).
fof(f3451, plain, ((e0 = op(e3, e3)) | (~ spl45_4 | ~ spl45_155)), inference(backward_demodulation, [], [f1411, f531])).
fof(f3435, plain, (~ spl45_21 | ~ spl45_22), inference(avatar_contradiction_clause, [], [f3434])).
fof(f3434, plain, ($false | (~ spl45_21 | ~ spl45_22)), inference(subsumption_resolution, [], [f3433, f242])).
fof(f3433, plain, ((e0 = e1) | (~ spl45_21 | ~ spl45_22)), inference(backward_demodulation, [], [f608, f604])).
fof(f3424, plain, (~ spl45_19 | ~ spl45_31), inference(avatar_split_clause, [], [f3422, f644, f593])).
fof(f3422, plain, (~ (e2 = op(e2, e3)) | ~ spl45_31), inference(backward_demodulation, [], [f232, f646])).
fof(f3418, plain, (spl45_19 | ~ spl45_40 | ~ spl45_102), inference(avatar_split_clause, [], [f3412, f1060, f682, f593])).
fof(f3412, plain, ((e2 = op(e2, e3)) | (~ spl45_40 | ~ spl45_102)), inference(backward_demodulation, [], [f1062, f684])).
fof(f3417, plain, (spl45_28 | ~ spl45_40 | ~ spl45_95), inference(avatar_split_clause, [], [f3411, f1027, f682, f631])).
fof(f3411, plain, ((e3 = op(e2, e1)) | (~ spl45_40 | ~ spl45_95)), inference(backward_demodulation, [], [f1029, f684])).
fof(f3392, plain, (~ spl45_31 | ~ spl45_55 | spl45_98), inference(avatar_split_clause, [], [f3389, f1041, f746, f644])).
fof(f3389, plain, (~ (e2 = op(e2, e0)) | (~ spl45_55 | spl45_98)), inference(backward_demodulation, [], [f1042, f748])).
fof(f3362, plain, (~ spl45_5 | ~ spl45_145 | spl45_148), inference(avatar_split_clause, [], [f3359, f1375, f1361, f534])).
fof(f3359, plain, (~ (e0 = op(e3, e2)) | (~ spl45_145 | spl45_148)), inference(backward_demodulation, [], [f1377, f1362])).
fof(f3336, plain, (~ spl45_2 | ~ spl45_10), inference(avatar_split_clause, [], [f3332, f555, f521])).
fof(f3332, plain, (~ (e1 = op(e3, e3)) | ~ spl45_10), inference(backward_demodulation, [], [f240, f557])).
fof(f3308, plain, (~ spl45_17 | ~ spl45_33), inference(avatar_split_clause, [], [f3301, f653, f585])).
fof(f3301, plain, (~ (e0 = op(e2, e3)) | ~ spl45_33), inference(backward_demodulation, [], [f215, f655])).
fof(f3300, plain, (spl45_38 | ~ spl45_39 | ~ spl45_94 | ~ spl45_95), inference(avatar_contradiction_clause, [], [f3299])).
fof(f3299, plain, ($false | (spl45_38 | ~ spl45_39 | ~ spl45_94 | ~ spl45_95)), inference(subsumption_resolution, [], [f3294, f675])).
fof(f675, plain, (~ (e1 = op(e1, e2)) | spl45_38), inference(avatar_component_clause, [], [f674])).
fof(f3294, plain, ((e1 = op(e1, e2)) | (~ spl45_39 | ~ spl45_94 | ~ spl45_95)), inference(backward_demodulation, [], [f3235, f680])).
fof(f3235, plain, ((e1 = op(e1, op(e1, e2))) | (~ spl45_94 | ~ spl45_95)), inference(forward_demodulation, [], [f1024, f1029])).
fof(f3297, plain, (spl45_27 | ~ spl45_39 | ~ spl45_95), inference(avatar_split_clause, [], [f3289, f1027, f678, f627])).
fof(f3289, plain, ((e2 = op(e2, e1)) | (~ spl45_39 | ~ spl45_95)), inference(backward_demodulation, [], [f1029, f680])).
fof(f3283, plain, (~ spl45_40 | ~ spl45_44), inference(avatar_split_clause, [], [f3279, f699, f682])).
fof(f3279, plain, (~ (e3 = op(e1, e2)) | ~ spl45_44), inference(backward_demodulation, [], [f227, f701])).
fof(f3276, plain, (~ spl45_46 | ~ spl45_57 | spl45_106), inference(avatar_contradiction_clause, [], [f3275])).
fof(f3275, plain, ($false | (~ spl45_46 | ~ spl45_57 | spl45_106)), inference(subsumption_resolution, [], [f3273, f757])).
fof(f3273, plain, (~ (e0 = op(e0, e1)) | (~ spl45_46 | spl45_106)), inference(backward_demodulation, [], [f1081, f710])).
fof(f1081, plain, (~ (e0 = op(e0, op(e1, e0))) | spl45_106), inference(avatar_component_clause, [], [f1080])).
fof(f3274, plain, (~ spl45_42 | ~ spl45_46), inference(avatar_split_clause, [], [f3268, f708, f691])).
fof(f3268, plain, (~ (e1 = op(e1, e1)) | ~ spl45_46), inference(backward_demodulation, [], [f224, f710])).
fof(f3266, plain, (spl45_27 | ~ spl45_54 | ~ spl45_111), inference(avatar_split_clause, [], [f3264, f1104, f742, f627])).
fof(f3264, plain, ((e2 = op(e2, e1)) | (~ spl45_54 | ~ spl45_111)), inference(backward_demodulation, [], [f1106, f744])).
fof(f3221, plain, (~ spl45_40 | spl45_146 | ~ spl45_147), inference(avatar_split_clause, [], [f3219, f1370, f1366, f682])).
fof(f3219, plain, (~ (e3 = op(e1, e2)) | (spl45_146 | ~ spl45_147)), inference(backward_demodulation, [], [f1368, f1371])).
fof(f3209, plain, (~ spl45_48 | ~ spl45_52 | ~ spl45_88), inference(avatar_split_clause, [], [f3208, f993, f733, f716])).
fof(f3208, plain, (~ (e3 = op(e1, e0)) | (~ spl45_52 | ~ spl45_88)), inference(forward_demodulation, [], [f2996, f735])).
fof(f2996, plain, (~ (op(e0, e3) = op(e1, e0)) | ~ spl45_88), inference(forward_demodulation, [], [f198, f995])).
fof(f3206, plain, (~ spl45_36 | ~ spl45_52), inference(avatar_split_clause, [], [f3205, f733, f665])).
fof(f3205, plain, (~ (e3 = op(e1, e3)) | ~ spl45_52), inference(forward_demodulation, [], [f212, f735])).
fof(f3196, plain, (~ spl45_1 | ~ spl45_9), inference(avatar_split_clause, [], [f3192, f551, f517])).
fof(f3192, plain, (~ (e0 = op(e3, e3)) | ~ spl45_9), inference(backward_demodulation, [], [f240, f553])).
fof(f3190, plain, (~ spl45_1 | ~ spl45_17), inference(avatar_split_clause, [], [f3186, f585, f517])).
fof(f3186, plain, (~ (e0 = op(e3, e3)) | ~ spl45_17), inference(backward_demodulation, [], [f217, f587])).
fof(f3108, plain, (spl45_12 | ~ spl45_18 | ~ spl45_90), inference(avatar_contradiction_clause, [], [f3107])).
fof(f3107, plain, ($false | (spl45_12 | ~ spl45_18 | ~ spl45_90)), inference(subsumption_resolution, [], [f3102, f564])).
fof(f3102, plain, ((e3 = op(e3, e1)) | (~ spl45_18 | ~ spl45_90)), inference(backward_demodulation, [], [f1004, f591])).
fof(f3096, plain, (~ spl45_9 | ~ spl45_25), inference(avatar_split_clause, [], [f3093, f619, f551])).
fof(f3093, plain, (~ (e0 = op(e3, e1)) | ~ spl45_25), inference(backward_demodulation, [], [f205, f621])).
fof(f3087, plain, (spl45_27 | ~ spl45_38 | ~ spl45_102), inference(avatar_contradiction_clause, [], [f3086])).
fof(f3086, plain, ($false | (spl45_27 | ~ spl45_38 | ~ spl45_102)), inference(subsumption_resolution, [], [f3085, f628])).
fof(f3085, plain, ((e2 = op(e2, e1)) | (~ spl45_38 | ~ spl45_102)), inference(forward_demodulation, [], [f1062, f676])).
fof(f3072, plain, (spl45_26 | ~ spl45_38 | ~ spl45_95), inference(avatar_split_clause, [], [f3071, f1027, f674, f623])).
fof(f3071, plain, ((e1 = op(e2, e1)) | (~ spl45_38 | ~ spl45_95)), inference(forward_demodulation, [], [f1029, f676])).
fof(f3070, plain, (~ spl45_25 | ~ spl45_141 | spl45_142), inference(avatar_split_clause, [], [f2904, f1345, f1340, f619])).
fof(f2904, plain, (~ (e0 = op(e2, e1)) | (~ spl45_141 | spl45_142)), inference(backward_demodulation, [], [f1347, f1341])).
fof(f1341, plain, ((e2 = op(e1, op(e1, e1))) | ~ spl45_141), inference(avatar_component_clause, [], [f1340])).
fof(f3069, plain, (~ spl45_19 | ~ spl45_35), inference(avatar_split_clause, [], [f3068, f661, f593])).
fof(f3068, plain, (~ (e2 = op(e2, e3)) | ~ spl45_35), inference(forward_demodulation, [], [f215, f663])).
fof(f3064, plain, (~ spl45_12 | ~ spl45_16), inference(avatar_split_clause, [], [f3063, f580, f563])).
fof(f3063, plain, (~ (e3 = op(e3, e1)) | ~ spl45_16), inference(forward_demodulation, [], [f236, f582])).
fof(f3062, plain, (~ spl45_49 | ~ spl45_52), inference(avatar_contradiction_clause, [], [f3061])).
fof(f3061, plain, ($false | (~ spl45_49 | ~ spl45_52)), inference(subsumption_resolution, [], [f3060, f244])).
fof(f244, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f3060, plain, ((e0 = e3) | (~ spl45_49 | ~ spl45_52)), inference(backward_demodulation, [], [f735, f723])).
fof(f3041, plain, (~ spl45_49 | ~ spl45_45 | ~ spl45_88), inference(avatar_split_clause, [], [f2997, f993, f704, f721])).
fof(f2997, plain, (~ (e0 = op(e0, e3)) | (~ spl45_45 | ~ spl45_88)), inference(forward_demodulation, [], [f2996, f706])).
fof(f3039, plain, (spl45_49 | ~ spl45_52 | ~ spl45_87 | ~ spl45_88), inference(avatar_split_clause, [], [f3030, f993, f988, f733, f721])).
fof(f3030, plain, ((e0 = op(e0, e3)) | (~ spl45_52 | ~ spl45_87 | ~ spl45_88)), inference(backward_demodulation, [], [f3017, f735])).
fof(f3017, plain, ((e0 = op(e0, op(e0, e3))) | (~ spl45_87 | ~ spl45_88)), inference(forward_demodulation, [], [f990, f995])).
fof(f3031, plain, (spl45_16 | ~ spl45_52 | ~ spl45_88), inference(avatar_split_clause, [], [f3027, f993, f733, f580])).
fof(f3027, plain, ((e3 = op(e3, e0)) | (~ spl45_52 | ~ spl45_88)), inference(backward_demodulation, [], [f995, f735])).
fof(f3026, plain, (~ spl45_14 | ~ spl45_62), inference(avatar_split_clause, [], [f3021, f776, f572])).
fof(f3021, plain, (~ (e1 = op(e3, e0)) | ~ spl45_62), inference(backward_demodulation, [], [f196, f778])).
fof(f2990, plain, (~ spl45_26 | spl45_42 | ~ spl45_94), inference(avatar_contradiction_clause, [], [f2989])).
fof(f2989, plain, ($false | (~ spl45_26 | spl45_42 | ~ spl45_94)), inference(subsumption_resolution, [], [f2988, f692])).
fof(f2988, plain, ((e1 = op(e1, e1)) | (~ spl45_26 | ~ spl45_94)), inference(forward_demodulation, [], [f1024, f625])).
fof(f2980, plain, (~ spl45_26 | ~ spl45_38 | ~ spl45_102), inference(avatar_contradiction_clause, [], [f2979])).
fof(f2979, plain, ($false | (~ spl45_26 | ~ spl45_38 | ~ spl45_102)), inference(subsumption_resolution, [], [f2978, f245])).
fof(f2978, plain, ((e1 = e2) | (~ spl45_26 | ~ spl45_38 | ~ spl45_102)), inference(forward_demodulation, [], [f2977, f625])).
fof(f2977, plain, ((e2 = op(e2, e1)) | (~ spl45_38 | ~ spl45_102)), inference(forward_demodulation, [], [f1062, f676])).
fof(f2976, plain, (~ spl45_7 | ~ spl45_17 | ~ spl45_82), inference(avatar_contradiction_clause, [], [f2975])).
fof(f2975, plain, ($false | (~ spl45_7 | ~ spl45_17 | ~ spl45_82)), inference(subsumption_resolution, [], [f2974, f243])).
fof(f2974, plain, ((e0 = e2) | (~ spl45_7 | ~ spl45_17 | ~ spl45_82)), inference(forward_demodulation, [], [f2973, f587])).
fof(f2973, plain, ((e2 = op(e2, e3)) | (~ spl45_7 | ~ spl45_82)), inference(forward_demodulation, [], [f967, f544])).
fof(f2972, plain, (~ spl45_14 | ~ spl45_17 | ~ spl45_90), inference(avatar_contradiction_clause, [], [f2971])).
fof(f2971, plain, ($false | (~ spl45_14 | ~ spl45_17 | ~ spl45_90)), inference(subsumption_resolution, [], [f2970, f246])).
fof(f2970, plain, ((e1 = e3) | (~ spl45_14 | ~ spl45_17 | ~ spl45_90)), inference(forward_demodulation, [], [f2969, f574])).
fof(f2969, plain, ((e3 = op(e3, e0)) | (~ spl45_17 | ~ spl45_90)), inference(forward_demodulation, [], [f1004, f587])).
fof(f2965, plain, (~ spl45_45 | spl45_61 | ~ spl45_106), inference(avatar_contradiction_clause, [], [f2964])).
fof(f2964, plain, ($false | (~ spl45_45 | spl45_61 | ~ spl45_106)), inference(subsumption_resolution, [], [f2963, f773])).
fof(f2963, plain, ((e0 = op(e0, e0)) | (~ spl45_45 | ~ spl45_106)), inference(forward_demodulation, [], [f1082, f706])).
fof(f2962, plain, (~ spl45_45 | spl45_57 | ~ spl45_107), inference(avatar_contradiction_clause, [], [f2961])).
fof(f2961, plain, ($false | (~ spl45_45 | spl45_57 | ~ spl45_107)), inference(subsumption_resolution, [], [f2960, f756])).
fof(f2960, plain, ((e0 = op(e0, e1)) | (~ spl45_45 | ~ spl45_107)), inference(forward_demodulation, [], [f1087, f706])).
fof(f2935, plain, (~ spl45_3 | ~ spl45_35), inference(avatar_split_clause, [], [f2931, f661, f525])).
fof(f2931, plain, (~ (e2 = op(e3, e3)) | ~ spl45_35), inference(backward_demodulation, [], [f216, f663])).
fof(f2915, plain, (~ spl45_14 | spl45_57 | ~ spl45_87), inference(avatar_contradiction_clause, [], [f2914])).
fof(f2914, plain, ($false | (~ spl45_14 | spl45_57 | ~ spl45_87)), inference(subsumption_resolution, [], [f2913, f756])).
fof(f2913, plain, ((e0 = op(e0, e1)) | (~ spl45_14 | ~ spl45_87)), inference(forward_demodulation, [], [f990, f574])).
fof(f2901, plain, (~ spl45_42 | ~ spl45_26), inference(avatar_split_clause, [], [f2757, f623, f691])).
fof(f2757, plain, (~ (e1 = op(e1, e1)) | ~ spl45_26), inference(forward_demodulation, [], [f203, f625])).
fof(f2884, plain, (~ spl45_35 | spl45_152 | ~ spl45_153), inference(avatar_split_clause, [], [f2882, f1400, f1396, f661])).
fof(f1396, plain, (spl45_152 <=> (e2 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl45_152])])).
fof(f1400, plain, (spl45_153 <=> (e1 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl45_153])])).
fof(f2882, plain, (~ (e2 = op(e1, e3)) | (spl45_152 | ~ spl45_153)), inference(backward_demodulation, [], [f1398, f1401])).
fof(f1401, plain, ((e1 = op(e3, op(e3, e3))) | ~ spl45_153), inference(avatar_component_clause, [], [f1400])).
fof(f1398, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | spl45_152), inference(avatar_component_clause, [], [f1396])).
fof(f2873, plain, (~ spl45_13 | ~ spl45_14), inference(avatar_contradiction_clause, [], [f2872])).
fof(f2872, plain, ($false | (~ spl45_13 | ~ spl45_14)), inference(subsumption_resolution, [], [f2871, f242])).
fof(f2871, plain, ((e0 = e1) | (~ spl45_13 | ~ spl45_14)), inference(forward_demodulation, [], [f574, f570])).
fof(f2870, plain, (~ spl45_38 | ~ spl45_39), inference(avatar_contradiction_clause, [], [f2869])).
fof(f2869, plain, ($false | (~ spl45_38 | ~ spl45_39)), inference(subsumption_resolution, [], [f2868, f245])).
fof(f2868, plain, ((e1 = e2) | (~ spl45_38 | ~ spl45_39)), inference(backward_demodulation, [], [f680, f676])).
fof(f2861, plain, (spl45_23 | ~ spl45_39 | ~ spl45_102), inference(avatar_contradiction_clause, [], [f2860])).
fof(f2860, plain, ($false | (spl45_23 | ~ spl45_39 | ~ spl45_102)), inference(subsumption_resolution, [], [f2859, f611])).
fof(f2859, plain, ((e2 = op(e2, e2)) | (~ spl45_39 | ~ spl45_102)), inference(forward_demodulation, [], [f1062, f680])).
fof(f2858, plain, (spl45_48 | ~ spl45_41 | ~ spl45_139), inference(avatar_split_clause, [], [f2857, f1331, f687, f716])).
fof(f2857, plain, ((e3 = op(e1, e0)) | (~ spl45_41 | ~ spl45_139)), inference(forward_demodulation, [], [f1332, f689])).
fof(f2852, plain, (~ spl45_47 | ~ spl45_31), inference(avatar_split_clause, [], [f2851, f644, f712])).
fof(f2851, plain, (~ (e2 = op(e1, e0)) | ~ spl45_31), inference(forward_demodulation, [], [f197, f646])).
fof(f2847, plain, (~ spl45_47 | ~ spl45_41 | spl45_141), inference(avatar_split_clause, [], [f2772, f1340, f687, f712])).
fof(f2772, plain, (~ (e2 = op(e1, e0)) | (~ spl45_41 | spl45_141)), inference(forward_demodulation, [], [f1342, f689])).
fof(f2846, plain, (~ spl45_47 | ~ spl45_50 | ~ spl45_64 | spl45_132), inference(avatar_split_clause, [], [f2784, f1297, f784, f725, f712])).
fof(f1297, plain, (spl45_132 <=> (e2 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl45_132])])).
fof(f2784, plain, (~ (e2 = op(e1, e0)) | (~ spl45_50 | ~ spl45_64 | spl45_132)), inference(forward_demodulation, [], [f2783, f727])).
fof(f2783, plain, (~ (e2 = op(op(e0, e3), e0)) | (~ spl45_64 | spl45_132)), inference(forward_demodulation, [], [f1299, f786])).
fof(f1299, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | spl45_132), inference(avatar_component_clause, [], [f1297])).
fof(f2813, plain, (~ spl45_46 | spl45_57 | ~ spl45_106), inference(avatar_contradiction_clause, [], [f2812])).
fof(f2812, plain, ($false | (~ spl45_46 | spl45_57 | ~ spl45_106)), inference(subsumption_resolution, [], [f2811, f756])).
fof(f2811, plain, ((e0 = op(e0, e1)) | (~ spl45_46 | ~ spl45_106)), inference(forward_demodulation, [], [f1082, f710])).
fof(f2756, plain, (spl45_5 | ~ spl45_3 | ~ spl45_155), inference(avatar_split_clause, [], [f2755, f1410, f525, f534])).
fof(f2755, plain, ((e0 = op(e3, e2)) | (~ spl45_3 | ~ spl45_155)), inference(backward_demodulation, [], [f1411, f527])).
fof(f2752, plain, (~ spl45_2 | ~ spl45_6), inference(avatar_split_clause, [], [f2748, f538, f521])).
fof(f2748, plain, (~ (e1 = op(e3, e3)) | ~ spl45_6), inference(backward_demodulation, [], [f241, f540])).
fof(f2736, plain, (spl45_12 | ~ spl45_36 | ~ spl45_85), inference(avatar_split_clause, [], [f2729, f979, f665, f563])).
fof(f2729, plain, ((e3 = op(e3, e1)) | (~ spl45_36 | ~ spl45_85)), inference(backward_demodulation, [], [f981, f667])).
fof(f2723, plain, (~ spl45_41 | ~ spl45_44), inference(avatar_contradiction_clause, [], [f2722])).
fof(f2722, plain, ($false | (~ spl45_41 | ~ spl45_44)), inference(subsumption_resolution, [], [f2721, f244])).
fof(f2721, plain, ((e0 = e3) | (~ spl45_41 | ~ spl45_44)), inference(backward_demodulation, [], [f701, f689])).
fof(f2697, plain, (~ spl45_63 | ~ spl45_64), inference(avatar_contradiction_clause, [], [f2696])).
fof(f2696, plain, ($false | (~ spl45_63 | ~ spl45_64)), inference(subsumption_resolution, [], [f2695, f247])).
fof(f2695, plain, ((e2 = e3) | (~ spl45_63 | ~ spl45_64)), inference(forward_demodulation, [], [f786, f782])).
fof(f2693, plain, (~ spl45_31 | spl45_55 | ~ spl45_98), inference(avatar_contradiction_clause, [], [f2692])).
fof(f2692, plain, ($false | (~ spl45_31 | spl45_55 | ~ spl45_98)), inference(subsumption_resolution, [], [f2691, f747])).
fof(f747, plain, (~ (e2 = op(e0, e2)) | spl45_55), inference(avatar_component_clause, [], [f746])).
fof(f2691, plain, ((e2 = op(e0, e2)) | (~ spl45_31 | ~ spl45_98)), inference(forward_demodulation, [], [f1043, f646])).
fof(f2661, plain, (~ spl45_54 | ~ spl45_17 | ~ spl45_24 | spl45_144), inference(avatar_split_clause, [], [f2613, f1357, f614, f585, f742])).
fof(f2613, plain, (~ (e1 = op(e0, e2)) | (~ spl45_17 | ~ spl45_24 | spl45_144)), inference(backward_demodulation, [], [f2598, f587])).
fof(f2598, plain, (~ (e1 = op(op(e2, e3), e2)) | (~ spl45_24 | spl45_144)), inference(backward_demodulation, [], [f1359, f616])).
fof(f2644, plain, (~ spl45_35 | ~ spl45_44 | spl45_141), inference(avatar_split_clause, [], [f2643, f1340, f699, f661])).
fof(f2643, plain, (~ (e2 = op(e1, e3)) | (~ spl45_44 | spl45_141)), inference(forward_demodulation, [], [f1342, f701])).
fof(f2627, plain, (~ spl45_11 | ~ spl45_59), inference(avatar_split_clause, [], [f2626, f763, f559])).
fof(f2626, plain, (~ (e2 = op(e3, e1)) | ~ spl45_59), inference(forward_demodulation, [], [f202, f765])).
fof(f2608, plain, (~ spl45_22 | ~ spl45_24), inference(avatar_contradiction_clause, [], [f2607])).
fof(f2607, plain, ($false | (~ spl45_22 | ~ spl45_24)), inference(subsumption_resolution, [], [f2606, f246])).
fof(f2606, plain, ((e1 = e3) | (~ spl45_22 | ~ spl45_24)), inference(backward_demodulation, [], [f616, f608])).
fof(f2595, plain, (~ spl45_25 | ~ spl45_46 | spl45_94), inference(avatar_contradiction_clause, [], [f2594])).
fof(f2594, plain, ($false | (~ spl45_25 | ~ spl45_46 | spl45_94)), inference(subsumption_resolution, [], [f2590, f710])).
fof(f2590, plain, (~ (e1 = op(e1, e0)) | (~ spl45_25 | spl45_94)), inference(backward_demodulation, [], [f1023, f621])).
fof(f1023, plain, (~ (e1 = op(e1, op(e2, e1))) | spl45_94), inference(avatar_component_clause, [], [f1022])).
fof(f2581, plain, (~ spl45_7 | ~ spl45_35 | ~ spl45_100), inference(avatar_contradiction_clause, [], [f2580])).
fof(f2580, plain, ($false | (~ spl45_7 | ~ spl45_35 | ~ spl45_100)), inference(subsumption_resolution, [], [f2579, f247])).
fof(f2579, plain, ((e2 = e3) | (~ spl45_7 | ~ spl45_35 | ~ spl45_100)), inference(forward_demodulation, [], [f2577, f544])).
fof(f2577, plain, ((e3 = op(e3, e2)) | (~ spl45_35 | ~ spl45_100)), inference(backward_demodulation, [], [f1052, f663])).
fof(f2567, plain, (~ spl45_42 | ~ spl45_44), inference(avatar_contradiction_clause, [], [f2566])).
fof(f2566, plain, ($false | (~ spl45_42 | ~ spl45_44)), inference(subsumption_resolution, [], [f2565, f246])).
fof(f2565, plain, ((e1 = e3) | (~ spl45_42 | ~ spl45_44)), inference(forward_demodulation, [], [f701, f693])).
fof(f2558, plain, (spl45_38 | ~ spl45_59 | ~ spl45_113), inference(avatar_contradiction_clause, [], [f2557])).
fof(f2557, plain, ($false | (spl45_38 | ~ spl45_59 | ~ spl45_113)), inference(subsumption_resolution, [], [f2555, f675])).
fof(f2555, plain, ((e1 = op(e1, e2)) | (~ spl45_59 | ~ spl45_113)), inference(backward_demodulation, [], [f1116, f765])).
fof(f2556, plain, (~ spl45_27 | ~ spl45_59), inference(avatar_split_clause, [], [f2553, f763, f627])).
fof(f2553, plain, (~ (e2 = op(e2, e1)) | ~ spl45_59), inference(backward_demodulation, [], [f201, f765])).
fof(f2552, plain, (~ spl45_61 | ~ spl45_63), inference(avatar_contradiction_clause, [], [f2551])).
fof(f2551, plain, ($false | (~ spl45_61 | ~ spl45_63)), inference(subsumption_resolution, [], [f2550, f243])).
fof(f2550, plain, ((e0 = e2) | (~ spl45_61 | ~ spl45_63)), inference(forward_demodulation, [], [f782, f774])).
fof(f2549, plain, (spl45_46 | ~ spl45_9 | ~ spl45_84), inference(avatar_split_clause, [], [f2548, f974, f551, f708])).
fof(f2548, plain, ((e1 = op(e1, e0)) | (~ spl45_9 | ~ spl45_84)), inference(forward_demodulation, [], [f976, f553])).
fof(f2523, plain, (~ spl45_48 | ~ spl45_16), inference(avatar_split_clause, [], [f2522, f580, f716])).
fof(f2522, plain, (~ (e3 = op(e1, e0)) | ~ spl45_16), inference(forward_demodulation, [], [f198, f582])).
fof(f2486, plain, (~ spl45_14 | ~ spl45_16), inference(avatar_contradiction_clause, [], [f2485])).
fof(f2485, plain, ($false | (~ spl45_14 | ~ spl45_16)), inference(subsumption_resolution, [], [f2484, f246])).
fof(f2484, plain, ((e1 = e3) | (~ spl45_14 | ~ spl45_16)), inference(forward_demodulation, [], [f582, f574])).
fof(f2483, plain, (~ spl45_21 | spl45_30 | ~ spl45_147), inference(avatar_contradiction_clause, [], [f2482])).
fof(f2482, plain, ($false | (~ spl45_21 | spl45_30 | ~ spl45_147)), inference(subsumption_resolution, [], [f2480, f641])).
fof(f641, plain, (~ (e1 = op(e2, e0)) | spl45_30), inference(avatar_component_clause, [], [f640])).
fof(f2480, plain, ((e1 = op(e2, e0)) | (~ spl45_21 | ~ spl45_147)), inference(backward_demodulation, [], [f1371, f604])).
fof(f2476, plain, (~ spl45_27 | ~ spl45_28), inference(avatar_contradiction_clause, [], [f2475])).
fof(f2475, plain, ($false | (~ spl45_27 | ~ spl45_28)), inference(subsumption_resolution, [], [f2474, f247])).
fof(f2474, plain, ((e2 = e3) | (~ spl45_27 | ~ spl45_28)), inference(forward_demodulation, [], [f633, f629])).
fof(f2473, plain, (~ spl45_1 | ~ spl45_33), inference(avatar_split_clause, [], [f2469, f653, f517])).
fof(f2469, plain, (~ (e0 = op(e3, e3)) | ~ spl45_33), inference(backward_demodulation, [], [f216, f655])).
fof(f2466, plain, (~ spl45_12 | ~ spl45_40 | ~ spl45_95), inference(avatar_split_clause, [], [f2463, f1027, f682, f563])).
fof(f2463, plain, (~ (e3 = op(e3, e1)) | (~ spl45_40 | ~ spl45_95)), inference(backward_demodulation, [], [f2247, f684])).
fof(f2247, plain, (~ (op(e1, e2) = op(e3, e1)) | ~ spl45_95), inference(forward_demodulation, [], [f205, f1029])).
fof(f2458, plain, (~ spl45_54 | ~ spl45_55), inference(avatar_contradiction_clause, [], [f2457])).
fof(f2457, plain, ($false | (~ spl45_54 | ~ spl45_55)), inference(subsumption_resolution, [], [f2456, f245])).
fof(f2456, plain, ((e1 = e2) | (~ spl45_54 | ~ spl45_55)), inference(forward_demodulation, [], [f748, f744])).
fof(f2441, plain, (~ spl45_47 | spl45_59 | ~ spl45_107), inference(avatar_contradiction_clause, [], [f2440])).
fof(f2440, plain, ($false | (~ spl45_47 | spl45_59 | ~ spl45_107)), inference(subsumption_resolution, [], [f2439, f764])).
fof(f764, plain, (~ (e2 = op(e0, e1)) | spl45_59), inference(avatar_component_clause, [], [f763])).
fof(f2439, plain, ((e2 = op(e0, e1)) | (~ spl45_47 | ~ spl45_107)), inference(forward_demodulation, [], [f1087, f714])).
fof(f2416, plain, (~ spl45_40 | ~ spl45_32 | ~ spl45_95), inference(avatar_split_clause, [], [f2415, f1027, f648, f682])).
fof(f2415, plain, (~ (e3 = op(e1, e2)) | (~ spl45_32 | ~ spl45_95)), inference(forward_demodulation, [], [f2252, f650])).
fof(f2252, plain, (~ (op(e1, e2) = op(e2, e0)) | ~ spl45_95), inference(forward_demodulation, [], [f230, f1029])).
fof(f2393, plain, (~ spl45_58 | ~ spl45_54), inference(avatar_split_clause, [], [f2392, f742, f759])).
fof(f2392, plain, (~ (e1 = op(e0, e1)) | ~ spl45_54), inference(forward_demodulation, [], [f221, f744])).
fof(f2386, plain, (spl45_39 | ~ spl45_27 | ~ spl45_95), inference(avatar_split_clause, [], [f2385, f1027, f627, f678])).
fof(f2385, plain, ((e2 = op(e1, e2)) | (~ spl45_27 | ~ spl45_95)), inference(forward_demodulation, [], [f1029, f629])).
fof(f2364, plain, (~ spl45_12 | spl45_34 | ~ spl45_84), inference(avatar_contradiction_clause, [], [f2363])).
fof(f2363, plain, ($false | (~ spl45_12 | spl45_34 | ~ spl45_84)), inference(subsumption_resolution, [], [f2362, f658])).
fof(f658, plain, (~ (e1 = op(e1, e3)) | spl45_34), inference(avatar_component_clause, [], [f657])).
fof(f2362, plain, ((e1 = op(e1, e3)) | (~ spl45_12 | ~ spl45_84)), inference(backward_demodulation, [], [f976, f565])).
fof(f2358, plain, (~ spl45_6 | ~ spl45_14), inference(avatar_split_clause, [], [f2354, f572, f538])).
fof(f2354, plain, (~ (e1 = op(e3, e2)) | ~ spl45_14), inference(backward_demodulation, [], [f237, f574])).
fof(f2352, plain, (~ spl45_25 | ~ spl45_27), inference(avatar_contradiction_clause, [], [f2351])).
fof(f2351, plain, ($false | (~ spl45_25 | ~ spl45_27)), inference(subsumption_resolution, [], [f2350, f243])).
fof(f2350, plain, ((e0 = e2) | (~ spl45_25 | ~ spl45_27)), inference(backward_demodulation, [], [f629, f621])).
fof(f2347, plain, (~ spl45_37 | spl45_46 | ~ spl45_94 | ~ spl45_95), inference(avatar_contradiction_clause, [], [f2346])).
fof(f2346, plain, ($false | (~ spl45_37 | spl45_46 | ~ spl45_94 | ~ spl45_95)), inference(subsumption_resolution, [], [f2343, f709])).
fof(f709, plain, (~ (e1 = op(e1, e0)) | spl45_46), inference(avatar_component_clause, [], [f708])).
fof(f2343, plain, ((e1 = op(e1, e0)) | (~ spl45_37 | ~ spl45_94 | ~ spl45_95)), inference(backward_demodulation, [], [f2244, f672])).
fof(f2244, plain, ((e1 = op(e1, op(e1, e2))) | (~ spl45_94 | ~ spl45_95)), inference(backward_demodulation, [], [f1024, f1029])).
fof(f2337, plain, (~ spl45_41 | ~ spl45_42 | spl45_143), inference(avatar_split_clause, [], [f2331, f1350, f691, f687])).
fof(f2331, plain, (~ (e0 = op(e1, e1)) | (~ spl45_42 | spl45_143)), inference(backward_demodulation, [], [f1352, f693])).
fof(f2336, plain, (~ spl45_42 | ~ spl45_141), inference(avatar_contradiction_clause, [], [f2335])).
fof(f2335, plain, ($false | (~ spl45_42 | ~ spl45_141)), inference(subsumption_resolution, [], [f2334, f245])).
fof(f2334, plain, ((e1 = e2) | (~ spl45_42 | ~ spl45_141)), inference(forward_demodulation, [], [f2330, f693])).
fof(f2330, plain, ((e2 = op(e1, e1)) | (~ spl45_42 | ~ spl45_141)), inference(backward_demodulation, [], [f1341, f693])).
fof(f2324, plain, (~ spl45_14 | ~ spl45_54 | ~ spl45_98), inference(avatar_split_clause, [], [f2318, f1041, f742, f572])).
fof(f2318, plain, (~ (e1 = op(e3, e0)) | (~ spl45_54 | ~ spl45_98)), inference(backward_demodulation, [], [f2261, f744])).
fof(f2261, plain, (~ (op(e0, e2) = op(e3, e0)) | ~ spl45_98), inference(forward_demodulation, [], [f199, f1043])).
fof(f2319, plain, (~ spl45_6 | ~ spl45_54), inference(avatar_split_clause, [], [f2312, f742, f538])).
fof(f2312, plain, (~ (e1 = op(e3, e2)) | ~ spl45_54), inference(backward_demodulation, [], [f208, f744])).
fof(f2311, plain, (~ spl45_61 | ~ spl45_64), inference(avatar_contradiction_clause, [], [f2310])).
fof(f2310, plain, ($false | (~ spl45_61 | ~ spl45_64)), inference(subsumption_resolution, [], [f2309, f244])).
fof(f2309, plain, ((e0 = e3) | (~ spl45_61 | ~ spl45_64)), inference(backward_demodulation, [], [f786, f774])).
fof(f2292, plain, (~ spl45_37 | ~ spl45_18 | ~ spl45_24 | spl45_148), inference(avatar_split_clause, [], [f2291, f1375, f614, f589, f670])).
fof(f2291, plain, (~ (e0 = op(e1, e2)) | (~ spl45_18 | ~ spl45_24 | spl45_148)), inference(forward_demodulation, [], [f2290, f591])).
fof(f2290, plain, (~ (e0 = op(op(e2, e3), e2)) | (~ spl45_24 | spl45_148)), inference(forward_demodulation, [], [f1377, f616])).
fof(f2286, plain, (~ spl45_14 | ~ spl45_133 | spl45_136), inference(avatar_split_clause, [], [f2285, f1315, f1301, f572])).
fof(f2285, plain, (~ (e1 = op(e3, e0)) | (~ spl45_133 | spl45_136)), inference(forward_demodulation, [], [f1317, f1302])).
fof(f1302, plain, ((e3 = op(e0, op(e0, e0))) | ~ spl45_133), inference(avatar_component_clause, [], [f1301])).
fof(f2274, plain, (~ spl45_43 | ~ spl45_47), inference(avatar_split_clause, [], [f2273, f712, f695])).
fof(f2273, plain, (~ (e2 = op(e1, e1)) | ~ spl45_47), inference(forward_demodulation, [], [f224, f714])).
fof(f2270, plain, (~ spl45_28 | spl45_140 | ~ spl45_141), inference(avatar_split_clause, [], [f2269, f1340, f1336, f631])).
fof(f2269, plain, (~ (e3 = op(e2, e1)) | (spl45_140 | ~ spl45_141)), inference(forward_demodulation, [], [f1338, f1341])).
fof(f2260, plain, (~ spl45_55 | ~ spl45_47 | ~ spl45_98), inference(avatar_split_clause, [], [f2259, f1041, f712, f746])).
fof(f2259, plain, (~ (e2 = op(e0, e2)) | (~ spl45_47 | ~ spl45_98)), inference(forward_demodulation, [], [f2258, f714])).
fof(f2258, plain, (~ (op(e0, e2) = op(e1, e0)) | ~ spl45_98), inference(forward_demodulation, [], [f197, f1043])).
fof(f2238, plain, (~ spl45_14 | ~ spl45_1 | spl45_153), inference(avatar_split_clause, [], [f2169, f1400, f517, f572])).
fof(f2169, plain, (~ (e1 = op(e3, e0)) | (~ spl45_1 | spl45_153)), inference(forward_demodulation, [], [f1402, f519])).
fof(f1402, plain, (~ (e1 = op(e3, op(e3, e3))) | spl45_153), inference(avatar_component_clause, [], [f1400])).
fof(f2234, plain, (spl45_6 | ~ spl45_18 | ~ spl45_82), inference(avatar_split_clause, [], [f2233, f965, f589, f538])).
fof(f2233, plain, ((e1 = op(e3, e2)) | (~ spl45_18 | ~ spl45_82)), inference(forward_demodulation, [], [f967, f591])).
fof(f2225, plain, (~ spl45_23 | ~ spl45_24), inference(avatar_contradiction_clause, [], [f2224])).
fof(f2224, plain, ($false | (~ spl45_23 | ~ spl45_24)), inference(subsumption_resolution, [], [f2223, f247])).
fof(f2223, plain, ((e2 = e3) | (~ spl45_23 | ~ spl45_24)), inference(backward_demodulation, [], [f616, f612])).
fof(f2217, plain, (~ spl45_26 | ~ spl45_28), inference(avatar_contradiction_clause, [], [f2216])).
fof(f2216, plain, ($false | (~ spl45_26 | ~ spl45_28)), inference(subsumption_resolution, [], [f2214, f246])).
fof(f2214, plain, ((e1 = e3) | (~ spl45_26 | ~ spl45_28)), inference(backward_demodulation, [], [f633, f625])).
fof(f2204, plain, (~ spl45_28 | ~ spl45_41 | ~ spl45_47 | spl45_140), inference(avatar_contradiction_clause, [], [f2203])).
fof(f2203, plain, ($false | (~ spl45_28 | ~ spl45_41 | ~ spl45_47 | spl45_140)), inference(subsumption_resolution, [], [f2199, f633])).
fof(f2199, plain, (~ (e3 = op(e2, e1)) | (~ spl45_41 | ~ spl45_47 | spl45_140)), inference(backward_demodulation, [], [f2085, f714])).
fof(f2085, plain, (~ (e3 = op(op(e1, e0), e1)) | (~ spl45_41 | spl45_140)), inference(forward_demodulation, [], [f1338, f689])).
fof(f2180, plain, (spl45_34 | ~ spl45_28 | ~ spl45_94), inference(avatar_split_clause, [], [f2179, f1022, f631, f657])).
fof(f2179, plain, ((e1 = op(e1, e3)) | (~ spl45_28 | ~ spl45_94)), inference(forward_demodulation, [], [f1024, f633])).
fof(f2178, plain, (spl45_40 | ~ spl45_28 | ~ spl45_95), inference(avatar_split_clause, [], [f2177, f1027, f631, f682])).
fof(f2177, plain, ((e3 = op(e1, e2)) | (~ spl45_28 | ~ spl45_95)), inference(forward_demodulation, [], [f1029, f633])).
fof(f2171, plain, (~ spl45_59 | ~ spl45_62 | spl45_135), inference(avatar_split_clause, [], [f2170, f1310, f776, f763])).
fof(f2170, plain, (~ (e2 = op(e0, e1)) | (~ spl45_62 | spl45_135)), inference(forward_demodulation, [], [f1312, f778])).
fof(f2164, plain, (~ spl45_38 | ~ spl45_6), inference(avatar_split_clause, [], [f2163, f538, f674])).
fof(f2163, plain, (~ (e1 = op(e1, e2)) | ~ spl45_6), inference(forward_demodulation, [], [f210, f540])).
fof(f2155, plain, (~ spl45_10 | ~ spl45_6), inference(avatar_split_clause, [], [f2054, f538, f555])).
fof(f2054, plain, (~ (e1 = op(e3, e1)) | ~ spl45_6), inference(forward_demodulation, [], [f239, f540])).
fof(f2154, plain, (~ spl45_1 | ~ spl45_3), inference(avatar_contradiction_clause, [], [f2153])).
fof(f2153, plain, ($false | (~ spl45_1 | ~ spl45_3)), inference(subsumption_resolution, [], [f2152, f243])).
fof(f2152, plain, ((e0 = e2) | (~ spl45_1 | ~ spl45_3)), inference(forward_demodulation, [], [f527, f519])).
fof(f2137, plain, (~ spl45_39 | ~ spl45_40), inference(avatar_contradiction_clause, [], [f2136])).
fof(f2136, plain, ($false | (~ spl45_39 | ~ spl45_40)), inference(subsumption_resolution, [], [f2135, f247])).
fof(f2135, plain, ((e2 = e3) | (~ spl45_39 | ~ spl45_40)), inference(backward_demodulation, [], [f684, f680])).
fof(f2121, plain, (~ spl45_51 | ~ spl45_59), inference(avatar_split_clause, [], [f2116, f763, f729])).
fof(f2116, plain, (~ (e2 = op(e0, e3)) | ~ spl45_59), inference(backward_demodulation, [], [f222, f765])).
fof(f2114, plain, (~ spl45_62 | ~ spl45_63), inference(avatar_contradiction_clause, [], [f2113])).
fof(f2113, plain, ($false | (~ spl45_62 | ~ spl45_63)), inference(subsumption_resolution, [], [f2112, f245])).
fof(f2112, plain, ((e1 = e2) | (~ spl45_62 | ~ spl45_63)), inference(forward_demodulation, [], [f782, f778])).
fof(f2060, plain, (~ spl45_18 | ~ spl45_34), inference(avatar_split_clause, [], [f2059, f657, f589])).
fof(f2059, plain, (~ (e1 = op(e2, e3)) | ~ spl45_34), inference(forward_demodulation, [], [f215, f659])).
fof(f2047, plain, (~ spl45_26 | ~ spl45_10), inference(avatar_split_clause, [], [f2046, f555, f623])).
fof(f2046, plain, (~ (e1 = op(e2, e1)) | ~ spl45_10), inference(forward_demodulation, [], [f205, f557])).
fof(f2031, plain, (~ spl45_10 | ~ spl45_11), inference(avatar_contradiction_clause, [], [f2030])).
fof(f2030, plain, ($false | (~ spl45_10 | ~ spl45_11)), inference(subsumption_resolution, [], [f2029, f245])).
fof(f2029, plain, ((e1 = e2) | (~ spl45_10 | ~ spl45_11)), inference(backward_demodulation, [], [f561, f557])).
fof(f2021, plain, (~ spl45_41 | spl45_45 | ~ spl45_143), inference(avatar_contradiction_clause, [], [f2020])).
fof(f2020, plain, ($false | (~ spl45_41 | spl45_45 | ~ spl45_143)), inference(subsumption_resolution, [], [f2018, f705])).
fof(f705, plain, (~ (e0 = op(e1, e0)) | spl45_45), inference(avatar_component_clause, [], [f704])).
fof(f2018, plain, ((e0 = op(e1, e0)) | (~ spl45_41 | ~ spl45_143)), inference(backward_demodulation, [], [f1351, f689])).
fof(f2019, plain, (~ spl45_33 | ~ spl45_41), inference(avatar_split_clause, [], [f2014, f687, f653])).
fof(f2014, plain, (~ (e0 = op(e1, e3)) | ~ spl45_41), inference(backward_demodulation, [], [f228, f689])).
fof(f1970, plain, (~ spl45_35 | ~ spl45_47), inference(avatar_split_clause, [], [f1969, f712, f661])).
fof(f1969, plain, (~ (e2 = op(e1, e3)) | ~ spl45_47), inference(forward_demodulation, [], [f226, f714])).
fof(f1968, plain, (~ spl45_35 | ~ spl45_51), inference(avatar_split_clause, [], [f1967, f729, f661])).
fof(f1967, plain, (~ (e2 = op(e1, e3)) | ~ spl45_51), inference(forward_demodulation, [], [f212, f731])).
fof(f1957, plain, (~ spl45_38 | ~ spl45_11 | spl45_84), inference(avatar_split_clause, [], [f1956, f974, f559, f674])).
fof(f1956, plain, (~ (e1 = op(e1, e2)) | (~ spl45_11 | spl45_84)), inference(forward_demodulation, [], [f975, f561])).
fof(f1942, plain, (~ spl45_9 | ~ spl45_11), inference(avatar_contradiction_clause, [], [f1941])).
fof(f1941, plain, ($false | (~ spl45_9 | ~ spl45_11)), inference(subsumption_resolution, [], [f1940, f243])).
fof(f1940, plain, ((e0 = e2) | (~ spl45_9 | ~ spl45_11)), inference(forward_demodulation, [], [f561, f553])).
fof(f1934, plain, (~ spl45_51 | ~ spl45_52), inference(avatar_contradiction_clause, [], [f1933])).
fof(f1933, plain, ($false | (~ spl45_51 | ~ spl45_52)), inference(subsumption_resolution, [], [f1932, f247])).
fof(f1932, plain, ((e2 = e3) | (~ spl45_51 | ~ spl45_52)), inference(backward_demodulation, [], [f735, f731])).
fof(f1919, plain, (~ spl45_59 | ~ spl45_33 | ~ spl45_44 | spl45_138), inference(avatar_split_clause, [], [f1918, f1327, f699, f653, f763])).
fof(f1918, plain, (~ (e2 = op(e0, e1)) | (~ spl45_33 | ~ spl45_44 | spl45_138)), inference(forward_demodulation, [], [f1917, f655])).
fof(f1917, plain, (~ (e2 = op(op(e1, e3), e1)) | (~ spl45_44 | spl45_138)), inference(forward_demodulation, [], [f1329, f701])).
fof(f1915, plain, (~ spl45_51 | spl45_152 | ~ spl45_155), inference(avatar_split_clause, [], [f1914, f1410, f1396, f729])).
fof(f1914, plain, (~ (e2 = op(e0, e3)) | (spl45_152 | ~ spl45_155)), inference(forward_demodulation, [], [f1398, f1411])).
fof(f1883, plain, (spl45_4 | ~ spl45_52 | ~ spl45_109), inference(avatar_contradiction_clause, [], [f1882])).
fof(f1882, plain, ($false | (spl45_4 | ~ spl45_52 | ~ spl45_109)), inference(subsumption_resolution, [], [f1881, f530])).
fof(f1881, plain, ((e3 = op(e3, e3)) | (~ spl45_52 | ~ spl45_109)), inference(forward_demodulation, [], [f1096, f735])).
fof(f1852, plain, (~ spl45_29 | ~ spl45_31), inference(avatar_contradiction_clause, [], [f1851])).
fof(f1851, plain, ($false | (~ spl45_29 | ~ spl45_31)), inference(subsumption_resolution, [], [f1850, f243])).
fof(f1850, plain, ((e0 = e2) | (~ spl45_29 | ~ spl45_31)), inference(backward_demodulation, [], [f646, f638])).
fof(f1838, plain, (~ spl45_43 | ~ spl45_44), inference(avatar_contradiction_clause, [], [f1837])).
fof(f1837, plain, ($false | (~ spl45_43 | ~ spl45_44)), inference(subsumption_resolution, [], [f1836, f247])).
fof(f1836, plain, ((e2 = e3) | (~ spl45_43 | ~ spl45_44)), inference(backward_demodulation, [], [f701, f697])).
fof(f1825, plain, (~ spl45_50 | ~ spl45_52), inference(avatar_contradiction_clause, [], [f1824])).
fof(f1824, plain, ($false | (~ spl45_50 | ~ spl45_52)), inference(subsumption_resolution, [], [f1823, f246])).
fof(f1823, plain, ((e1 = e3) | (~ spl45_50 | ~ spl45_52)), inference(forward_demodulation, [], [f735, f727])).
fof(f1822, plain, (~ spl45_53 | ~ spl45_55), inference(avatar_contradiction_clause, [], [f1821])).
fof(f1821, plain, ($false | (~ spl45_53 | ~ spl45_55)), inference(subsumption_resolution, [], [f1820, f243])).
fof(f1820, plain, ((e0 = e2) | (~ spl45_53 | ~ spl45_55)), inference(backward_demodulation, [], [f748, f740])).
fof(f1794, plain, (~ spl45_6 | ~ spl45_3 | spl45_153), inference(avatar_split_clause, [], [f1793, f1400, f525, f538])).
fof(f1793, plain, (~ (e1 = op(e3, e2)) | (~ spl45_3 | spl45_153)), inference(forward_demodulation, [], [f1402, f527])).
fof(f1787, plain, (~ spl45_59 | ~ spl45_55), inference(avatar_split_clause, [], [f1786, f746, f763])).
fof(f1786, plain, (~ (e2 = op(e0, e1)) | ~ spl45_55), inference(forward_demodulation, [], [f221, f748])).
fof(f1772, plain, (~ spl45_6 | ~ spl45_8), inference(avatar_contradiction_clause, [], [f1771])).
fof(f1771, plain, ($false | (~ spl45_6 | ~ spl45_8)), inference(subsumption_resolution, [], [f1770, f246])).
fof(f1770, plain, ((e1 = e3) | (~ spl45_6 | ~ spl45_8)), inference(backward_demodulation, [], [f548, f540])).
fof(f1767, plain, (~ spl45_14 | ~ spl45_30), inference(avatar_split_clause, [], [f1765, f640, f572])).
fof(f1765, plain, (~ (e1 = op(e3, e0)) | ~ spl45_30), inference(backward_demodulation, [], [f199, f642])).
fof(f1760, plain, (~ spl45_46 | ~ spl45_47), inference(avatar_contradiction_clause, [], [f1759])).
fof(f1759, plain, ($false | (~ spl45_46 | ~ spl45_47)), inference(subsumption_resolution, [], [f1758, f245])).
fof(f1758, plain, ((e1 = e2) | (~ spl45_46 | ~ spl45_47)), inference(backward_demodulation, [], [f714, f710])).
fof(f1757, plain, (~ spl45_47 | ~ spl45_48), inference(avatar_contradiction_clause, [], [f1756])).
fof(f1756, plain, ($false | (~ spl45_47 | ~ spl45_48)), inference(subsumption_resolution, [], [f1755, f247])).
fof(f1755, plain, ((e2 = e3) | (~ spl45_47 | ~ spl45_48)), inference(forward_demodulation, [], [f718, f714])).
fof(f1738, plain, (~ spl45_32 | ~ spl45_20), inference(avatar_split_clause, [], [f1737, f597, f648])).
fof(f1737, plain, (~ (e3 = op(e2, e0)) | ~ spl45_20), inference(forward_demodulation, [], [f232, f599])).
fof(f1736, plain, (~ spl45_32 | ~ spl45_21 | spl45_145), inference(avatar_split_clause, [], [f1720, f1361, f602, f648])).
fof(f1720, plain, (~ (e3 = op(e2, e0)) | (~ spl45_21 | spl45_145)), inference(backward_demodulation, [], [f1363, f604])).
fof(f1735, plain, (~ spl45_30 | ~ spl45_21 | spl45_147), inference(avatar_split_clause, [], [f1721, f1370, f602, f640])).
fof(f1721, plain, (~ (e1 = op(e2, e0)) | (~ spl45_21 | spl45_147)), inference(backward_demodulation, [], [f1372, f604])).
fof(f1716, plain, (~ spl45_26 | ~ spl45_27), inference(avatar_contradiction_clause, [], [f1715])).
fof(f1715, plain, ($false | (~ spl45_26 | ~ spl45_27)), inference(subsumption_resolution, [], [f1714, f245])).
fof(f1714, plain, ((e1 = e2) | (~ spl45_26 | ~ spl45_27)), inference(backward_demodulation, [], [f629, f625])).
fof(f1713, plain, (~ spl45_33 | ~ spl45_36), inference(avatar_contradiction_clause, [], [f1712])).
fof(f1712, plain, ($false | (~ spl45_33 | ~ spl45_36)), inference(subsumption_resolution, [], [f1710, f244])).
fof(f1710, plain, ((e0 = e3) | (~ spl45_33 | ~ spl45_36)), inference(backward_demodulation, [], [f667, f655])).
fof(f1666, plain, (~ spl45_2 | ~ spl45_3), inference(avatar_contradiction_clause, [], [f1665])).
fof(f1665, plain, ($false | (~ spl45_2 | ~ spl45_3)), inference(subsumption_resolution, [], [f1664, f245])).
fof(f1664, plain, ((e1 = e2) | (~ spl45_2 | ~ spl45_3)), inference(forward_demodulation, [], [f527, f523])).
fof(f1663, plain, (~ spl45_5 | ~ spl45_8), inference(avatar_contradiction_clause, [], [f1662])).
fof(f1662, plain, ($false | (~ spl45_5 | ~ spl45_8)), inference(subsumption_resolution, [], [f1661, f244])).
fof(f1661, plain, ((e0 = e3) | (~ spl45_5 | ~ spl45_8)), inference(backward_demodulation, [], [f548, f536])).
fof(f1635, plain, (~ spl45_10 | ~ spl45_41 | ~ spl45_84), inference(avatar_contradiction_clause, [], [f1634])).
fof(f1634, plain, ($false | (~ spl45_10 | ~ spl45_41 | ~ spl45_84)), inference(subsumption_resolution, [], [f1633, f242])).
fof(f1633, plain, ((e0 = e1) | (~ spl45_10 | ~ spl45_41 | ~ spl45_84)), inference(forward_demodulation, [], [f1632, f689])).
fof(f1632, plain, ((e1 = op(e1, e1)) | (~ spl45_10 | ~ spl45_84)), inference(forward_demodulation, [], [f976, f557])).
fof(f1631, plain, (~ spl45_10 | spl45_34 | ~ spl45_85), inference(avatar_contradiction_clause, [], [f1630])).
fof(f1630, plain, ($false | (~ spl45_10 | spl45_34 | ~ spl45_85)), inference(subsumption_resolution, [], [f1629, f658])).
fof(f1629, plain, ((e1 = op(e1, e3)) | (~ spl45_10 | ~ spl45_85)), inference(forward_demodulation, [], [f981, f557])).
fof(f1628, plain, (~ spl45_57 | ~ spl45_41), inference(avatar_split_clause, [], [f1627, f687, f755])).
fof(f1627, plain, (~ (e0 = op(e0, e1)) | ~ spl45_41), inference(forward_demodulation, [], [f200, f689])).
fof(f1624, plain, (~ spl45_58 | ~ spl45_10), inference(avatar_split_clause, [], [f1623, f555, f759])).
fof(f1623, plain, (~ (e1 = op(e0, e1)) | ~ spl45_10), inference(forward_demodulation, [], [f202, f557])).
fof(f1618, plain, (~ spl45_56 | ~ spl45_8), inference(avatar_split_clause, [], [f1617, f546, f750])).
fof(f1617, plain, (~ (e3 = op(e0, e2)) | ~ spl45_8), inference(forward_demodulation, [], [f208, f548])).
fof(f1603, plain, (~ spl45_48 | ~ spl45_41 | spl45_139), inference(avatar_split_clause, [], [f1479, f1331, f687, f716])).
fof(f1479, plain, (~ (e3 = op(e1, e0)) | (~ spl45_41 | spl45_139)), inference(backward_demodulation, [], [f1333, f689])).
fof(f1587, plain, (~ spl45_21 | ~ spl45_17), inference(avatar_split_clause, [], [f1586, f585, f602])).
fof(f1586, plain, (~ (e0 = op(e2, e2)) | ~ spl45_17), inference(forward_demodulation, [], [f235, f587])).
fof(f1584, plain, (~ spl45_17 | ~ spl45_20), inference(avatar_contradiction_clause, [], [f1583])).
fof(f1583, plain, ($false | (~ spl45_17 | ~ spl45_20)), inference(subsumption_resolution, [], [f1582, f244])).
fof(f1582, plain, ((e0 = e3) | (~ spl45_17 | ~ spl45_20)), inference(forward_demodulation, [], [f599, f587])).
fof(f1571, plain, (~ spl45_9 | ~ spl45_2 | spl45_155), inference(avatar_split_clause, [], [f1567, f1410, f521, f551])).
fof(f1567, plain, (~ (e0 = op(e3, e1)) | (~ spl45_2 | spl45_155)), inference(backward_demodulation, [], [f1412, f523])).
fof(f1564, plain, (~ spl45_7 | ~ spl45_8), inference(avatar_contradiction_clause, [], [f1563])).
fof(f1563, plain, ($false | (~ spl45_7 | ~ spl45_8)), inference(subsumption_resolution, [], [f1562, f247])).
fof(f1562, plain, ((e2 = e3) | (~ spl45_7 | ~ spl45_8)), inference(backward_demodulation, [], [f548, f544])).
fof(f1561, plain, (~ spl45_9 | ~ spl45_10), inference(avatar_contradiction_clause, [], [f1560])).
fof(f1560, plain, ($false | (~ spl45_9 | ~ spl45_10)), inference(subsumption_resolution, [], [f1559, f242])).
fof(f1559, plain, ((e0 = e1) | (~ spl45_9 | ~ spl45_10)), inference(backward_demodulation, [], [f557, f553])).
fof(f1558, plain, (~ spl45_10 | ~ spl45_12), inference(avatar_contradiction_clause, [], [f1557])).
fof(f1557, plain, ($false | (~ spl45_10 | ~ spl45_12)), inference(subsumption_resolution, [], [f1556, f246])).
fof(f1556, plain, ((e1 = e3) | (~ spl45_10 | ~ spl45_12)), inference(backward_demodulation, [], [f565, f557])).
fof(f1552, plain, (~ spl45_14 | ~ spl45_15), inference(avatar_contradiction_clause, [], [f1551])).
fof(f1551, plain, ($false | (~ spl45_14 | ~ spl45_15)), inference(subsumption_resolution, [], [f1550, f245])).
fof(f1550, plain, ((e1 = e2) | (~ spl45_14 | ~ spl45_15)), inference(backward_demodulation, [], [f578, f574])).
fof(f1549, plain, (~ spl45_15 | ~ spl45_16), inference(avatar_contradiction_clause, [], [f1548])).
fof(f1548, plain, ($false | (~ spl45_15 | ~ spl45_16)), inference(subsumption_resolution, [], [f1547, f247])).
fof(f1547, plain, ((e2 = e3) | (~ spl45_15 | ~ spl45_16)), inference(backward_demodulation, [], [f582, f578])).
fof(f1546, plain, (~ spl45_4 | ~ spl45_16), inference(avatar_split_clause, [], [f1544, f580, f529])).
fof(f1544, plain, (~ (e3 = op(e3, e3)) | ~ spl45_16), inference(backward_demodulation, [], [f238, f582])).
fof(f1542, plain, (~ spl45_17 | ~ spl45_18), inference(avatar_contradiction_clause, [], [f1541])).
fof(f1541, plain, ($false | (~ spl45_17 | ~ spl45_18)), inference(subsumption_resolution, [], [f1540, f242])).
fof(f1540, plain, ((e0 = e1) | (~ spl45_17 | ~ spl45_18)), inference(backward_demodulation, [], [f591, f587])).
fof(f1539, plain, (~ spl45_18 | ~ spl45_20), inference(avatar_contradiction_clause, [], [f1538])).
fof(f1538, plain, ($false | (~ spl45_18 | ~ spl45_20)), inference(subsumption_resolution, [], [f1537, f246])).
fof(f1537, plain, ((e1 = e3) | (~ spl45_18 | ~ spl45_20)), inference(backward_demodulation, [], [f599, f591])).
fof(f1535, plain, (spl45_19 | ~ spl45_20 | ~ spl45_81 | ~ spl45_82), inference(avatar_split_clause, [], [f1530, f965, f960, f597, f593])).
fof(f1530, plain, ((e2 = op(e2, e3)) | (~ spl45_20 | ~ spl45_81 | ~ spl45_82)), inference(backward_demodulation, [], [f1419, f599])).
fof(f1419, plain, ((e2 = op(e2, op(e2, e3))) | (~ spl45_81 | ~ spl45_82)), inference(forward_demodulation, [], [f962, f967])).
fof(f1534, plain, (~ spl45_12 | ~ spl45_20 | ~ spl45_82), inference(avatar_split_clause, [], [f1529, f965, f597, f563])).
fof(f1529, plain, (~ (e3 = op(e3, e1)) | (~ spl45_20 | ~ spl45_82)), inference(backward_demodulation, [], [f1417, f599])).
fof(f1417, plain, (~ (op(e2, e3) = op(e3, e1)) | ~ spl45_82), inference(backward_demodulation, [], [f239, f967])).
fof(f1513, plain, (~ spl45_23 | ~ spl45_27), inference(avatar_split_clause, [], [f1510, f627, f610])).
fof(f1510, plain, (~ (e2 = op(e2, e2)) | ~ spl45_27), inference(backward_demodulation, [], [f233, f629])).
fof(f1495, plain, (~ spl45_38 | ~ spl45_40), inference(avatar_contradiction_clause, [], [f1494])).
fof(f1494, plain, ($false | (~ spl45_38 | ~ spl45_40)), inference(subsumption_resolution, [], [f1493, f246])).
fof(f1493, plain, ((e1 = e3) | (~ spl45_38 | ~ spl45_40)), inference(backward_demodulation, [], [f684, f676])).
fof(f1481, plain, (~ spl45_25 | ~ spl45_41), inference(avatar_split_clause, [], [f1472, f687, f619])).
fof(f1472, plain, (~ (e0 = op(e2, e1)) | ~ spl45_41), inference(backward_demodulation, [], [f203, f689])).
fof(f1470, plain, (~ spl45_37 | ~ spl45_45), inference(avatar_split_clause, [], [f1465, f704, f670])).
fof(f1465, plain, (~ (e0 = op(e1, e2)) | ~ spl45_45), inference(backward_demodulation, [], [f225, f706])).
fof(f1461, plain, (~ spl45_50 | ~ spl45_51), inference(avatar_contradiction_clause, [], [f1460])).
fof(f1460, plain, ($false | (~ spl45_50 | ~ spl45_51)), inference(subsumption_resolution, [], [f1459, f245])).
fof(f1459, plain, ((e1 = e2) | (~ spl45_50 | ~ spl45_51)), inference(backward_demodulation, [], [f731, f727])).
fof(f1433, plain, (~ spl45_53 | ~ spl45_61), inference(avatar_split_clause, [], [f1424, f772, f738])).
fof(f1424, plain, (~ (e0 = op(e0, e2)) | ~ spl45_61), inference(backward_demodulation, [], [f219, f774])).
fof(f1432, plain, (~ spl45_57 | ~ spl45_61), inference(avatar_split_clause, [], [f1423, f772, f755])).
fof(f1423, plain, (~ (e0 = op(e0, e1)) | ~ spl45_61), inference(backward_demodulation, [], [f218, f774])).
fof(f1431, plain, (~ spl45_13 | ~ spl45_61), inference(avatar_split_clause, [], [f1422, f772, f568])).
fof(f1422, plain, (~ (e0 = op(e3, e0)) | ~ spl45_61), inference(backward_demodulation, [], [f196, f774])).
fof(f1415, plain, (~ spl45_3 | ~ spl45_150 | ~ spl45_155), inference(avatar_split_clause, [], [f495, f1410, f1387, f525])).
fof(f495, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3))), inference(cnf_transformation, [], [f54])).
fof(f54, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax29)).
fof(f1414, plain, (~ spl45_3 | ~ spl45_154 | ~ spl45_153), inference(avatar_split_clause, [], [f494, f1400, f1405, f525])).
fof(f494, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e1 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax28)).
fof(f1413, plain, (~ spl45_2 | ~ spl45_152 | ~ spl45_155), inference(avatar_split_clause, [], [f493, f1410, f1396, f521])).
fof(f493, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e2 = op(op(e3, op(e3, e3)), e3)) & (e1 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax27)).
fof(f1408, plain, (~ spl45_2 | ~ spl45_154 | ~ spl45_151), inference(avatar_split_clause, [], [f492, f1391, f1405, f521])).
fof(f492, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), e3)) & (e1 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax26)).
fof(f1403, plain, (~ spl45_1 | ~ spl45_152 | ~ spl45_153), inference(avatar_split_clause, [], [f491, f1400, f1396, f517])).
fof(f491, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e1 = op(e3, op(e3, e3))) & (e2 = op(op(e3, op(e3, e3)), e3)) & (e0 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax25)).
fof(f1385, plain, (~ spl45_24 | ~ spl45_144 | ~ spl45_149), inference(avatar_split_clause, [], [f489, f1380, f1357, f614])).
fof(f489, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax23)).
fof(f1384, plain, (~ spl45_24 | ~ spl45_148 | ~ spl45_147), inference(avatar_split_clause, [], [f488, f1370, f1375, f614])).
fof(f488, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax22)).
fof(f1383, plain, (~ spl45_22 | ~ spl45_146 | ~ spl45_149), inference(avatar_split_clause, [], [f487, f1380, f1366, f606])).
fof(f487, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e3 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax21)).
fof(f1378, plain, (~ spl45_22 | ~ spl45_148 | ~ spl45_145), inference(avatar_split_clause, [], [f486, f1361, f1375, f606])).
fof(f486, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax20)).
fof(f1373, plain, (~ spl45_21 | ~ spl45_146 | ~ spl45_147), inference(avatar_split_clause, [], [f485, f1370, f1366, f602])).
fof(f485, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(e2, op(e2, e2))) & (e3 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax19)).
fof(f1364, plain, (~ spl45_21 | ~ spl45_144 | ~ spl45_145), inference(avatar_split_clause, [], [f484, f1361, f1357, f602])).
fof(f484, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax18)).
fof(f1355, plain, (~ spl45_44 | ~ spl45_138 | ~ spl45_143), inference(avatar_split_clause, [], [f483, f1350, f1327, f699])).
fof(f483, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax17)).
fof(f1354, plain, (~ spl45_44 | ~ spl45_142 | ~ spl45_141), inference(avatar_split_clause, [], [f482, f1340, f1345, f699])).
fof(f482, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax16)).
fof(f1353, plain, (~ spl45_43 | ~ spl45_140 | ~ spl45_143), inference(avatar_split_clause, [], [f481, f1350, f1336, f695])).
fof(f481, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax15)).
fof(f1348, plain, (~ spl45_43 | ~ spl45_142 | ~ spl45_139), inference(avatar_split_clause, [], [f480, f1331, f1345, f695])).
fof(f480, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax14)).
fof(f1343, plain, (~ spl45_41 | ~ spl45_140 | ~ spl45_141), inference(avatar_split_clause, [], [f479, f1340, f1336, f687])).
fof(f479, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax13)).
fof(f1325, plain, (~ spl45_64 | ~ spl45_132 | ~ spl45_137), inference(avatar_split_clause, [], [f477, f1320, f1297, f784])).
fof(f477, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3)), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3)), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax11)).
fof(f1324, plain, (~ spl45_64 | ~ spl45_136 | ~ spl45_135), inference(avatar_split_clause, [], [f476, f1310, f1315, f784])).
fof(f476, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3)), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3)), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax10)).
fof(f1318, plain, (~ spl45_63 | ~ spl45_136 | ~ spl45_133), inference(avatar_split_clause, [], [f474, f1301, f1315, f780])).
fof(f474, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2)), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2)), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e2)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax8)).
fof(f1295, plain, (spl45_130 | spl45_129 | spl45_128 | spl45_127 | spl45_126 | spl45_125 | spl45_124 | spl45_123 | spl45_122 | spl45_121 | spl45_120 | spl45_119 | spl45_118 | spl45_117 | spl45_116 | spl45_4), inference(avatar_split_clause, [], [f458, f529, f1130, f1140, f1150, f1160, f1170, f1180, f1190, f1200, f1210, f1220, f1230, f1240, f1250, f1260, f1270])).
fof(f1270, plain, (spl45_130 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl45_130])])).
fof(f1260, plain, (spl45_129 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl45_129])])).
fof(f1250, plain, (spl45_128 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl45_128])])).
fof(f1240, plain, (spl45_127 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl45_127])])).
fof(f1230, plain, (spl45_126 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl45_126])])).
fof(f1220, plain, (spl45_125 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl45_125])])).
fof(f1210, plain, (spl45_124 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl45_124])])).
fof(f1200, plain, (spl45_123 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl45_123])])).
fof(f1190, plain, (spl45_122 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl45_122])])).
fof(f1180, plain, (spl45_121 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl45_121])])).
fof(f1170, plain, (spl45_120 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl45_120])])).
fof(f1160, plain, (spl45_119 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl45_119])])).
fof(f1150, plain, (spl45_118 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl45_118])])).
fof(f1140, plain, (spl45_117 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl45_117])])).
fof(f1130, plain, (spl45_116 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl45_116])])).
fof(f458, plain, ((e3 = op(e3, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((((((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e2 = op(e3, e3))) & (~ (e1 = op(e3, e3)) | ~ (e1 = op(e3, e3))) & (~ (e0 = op(e3, e3)) | ~ (e0 = op(e3, e3)))) | ~ (e3 = op(e3, e3))) & (e3 = op(e3, e3))) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30) & (((e3 = op(e3, op(e3, e3))) & ~ (op(e3, e3) = op(e3, e3)) & (op(e3, e3) = op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15) & ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f30, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55])).
fof(f55, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e55])).
fof(e55, plain, (sP0 <=> (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f56, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP1), inference(usedef, [], [e56])).
fof(e56, plain, (sP1 <=> (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f57, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP2), inference(usedef, [], [e57])).
fof(e57, plain, (sP2 <=> (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f58, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP3), inference(usedef, [], [e58])).
fof(e58, plain, (sP3 <=> (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f59, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP4), inference(usedef, [], [e59])).
fof(e59, plain, (sP4 <=> (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f60, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP5), inference(usedef, [], [e60])).
fof(e60, plain, (sP5 <=> (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f61, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP6), inference(usedef, [], [e61])).
fof(e61, plain, (sP6 <=> (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f62, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP7), inference(usedef, [], [e62])).
fof(e62, plain, (sP7 <=> (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f63, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP8), inference(usedef, [], [e63])).
fof(e63, plain, (sP8 <=> (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f64, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP9), inference(usedef, [], [e64])).
fof(e64, plain, (sP9 <=> (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f65, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP10), inference(usedef, [], [e65])).
fof(e65, plain, (sP10 <=> (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f66, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP11), inference(usedef, [], [e66])).
fof(e66, plain, (sP11 <=> (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f67, plain, ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP12), inference(usedef, [], [e67])).
fof(e67, plain, (sP12 <=> (~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f68, plain, ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP13), inference(usedef, [], [e68])).
fof(e68, plain, (sP13 <=> (~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f69, plain, ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP14), inference(usedef, [], [e69])).
fof(e69, plain, (sP14 <=> (~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f70, plain, (((e0 = op(e0, op(e0, e0))) & ~ (op(e0, e0) = op(e0, e0)) & (op(e0, e0) = op(e0, e0))) | ~ sP15), inference(usedef, [], [e70])).
fof(e70, plain, (sP15 <=> ((e0 = op(e0, op(e0, e0))) & ~ (op(e0, e0) = op(e0, e0)) & (op(e0, e0) = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f71, plain, (((e1 = op(e1, op(e0, e1))) & ~ (op(e0, e1) = op(e1, e1)) & (op(e0, e1) = op(e1, e0))) | ~ sP16), inference(usedef, [], [e71])).
fof(e71, plain, (sP16 <=> ((e1 = op(e1, op(e0, e1))) & ~ (op(e0, e1) = op(e1, e1)) & (op(e0, e1) = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f72, plain, (((e2 = op(e2, op(e0, e2))) & ~ (op(e0, e2) = op(e2, e2)) & (op(e0, e2) = op(e2, e0))) | ~ sP17), inference(usedef, [], [e72])).
fof(e72, plain, (sP17 <=> ((e2 = op(e2, op(e0, e2))) & ~ (op(e0, e2) = op(e2, e2)) & (op(e0, e2) = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f73, plain, (((e3 = op(e3, op(e0, e3))) & ~ (op(e0, e3) = op(e3, e3)) & (op(e0, e3) = op(e3, e0))) | ~ sP18), inference(usedef, [], [e73])).
fof(e73, plain, (sP18 <=> ((e3 = op(e3, op(e0, e3))) & ~ (op(e0, e3) = op(e3, e3)) & (op(e0, e3) = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f74, plain, (((e0 = op(e0, op(e1, e0))) & ~ (op(e0, e0) = op(e1, e0)) & (op(e0, e1) = op(e1, e0))) | ~ sP19), inference(usedef, [], [e74])).
fof(e74, plain, (sP19 <=> ((e0 = op(e0, op(e1, e0))) & ~ (op(e0, e0) = op(e1, e0)) & (op(e0, e1) = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f75, plain, (((e1 = op(e1, op(e1, e1))) & ~ (op(e1, e1) = op(e1, e1)) & (op(e1, e1) = op(e1, e1))) | ~ sP20), inference(usedef, [], [e75])).
fof(e75, plain, (sP20 <=> ((e1 = op(e1, op(e1, e1))) & ~ (op(e1, e1) = op(e1, e1)) & (op(e1, e1) = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f76, plain, (((e2 = op(e2, op(e1, e2))) & ~ (op(e1, e2) = op(e2, e2)) & (op(e1, e2) = op(e2, e1))) | ~ sP21), inference(usedef, [], [e76])).
fof(e76, plain, (sP21 <=> ((e2 = op(e2, op(e1, e2))) & ~ (op(e1, e2) = op(e2, e2)) & (op(e1, e2) = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f77, plain, (((e3 = op(e3, op(e1, e3))) & ~ (op(e1, e3) = op(e3, e3)) & (op(e1, e3) = op(e3, e1))) | ~ sP22), inference(usedef, [], [e77])).
fof(e77, plain, (sP22 <=> ((e3 = op(e3, op(e1, e3))) & ~ (op(e1, e3) = op(e3, e3)) & (op(e1, e3) = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f78, plain, (((e0 = op(e0, op(e2, e0))) & ~ (op(e0, e0) = op(e2, e0)) & (op(e0, e2) = op(e2, e0))) | ~ sP23), inference(usedef, [], [e78])).
fof(e78, plain, (sP23 <=> ((e0 = op(e0, op(e2, e0))) & ~ (op(e0, e0) = op(e2, e0)) & (op(e0, e2) = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f79, plain, (((e1 = op(e1, op(e2, e1))) & ~ (op(e1, e1) = op(e2, e1)) & (op(e1, e2) = op(e2, e1))) | ~ sP24), inference(usedef, [], [e79])).
fof(e79, plain, (sP24 <=> ((e1 = op(e1, op(e2, e1))) & ~ (op(e1, e1) = op(e2, e1)) & (op(e1, e2) = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f80, plain, (((e2 = op(e2, op(e2, e2))) & ~ (op(e2, e2) = op(e2, e2)) & (op(e2, e2) = op(e2, e2))) | ~ sP25), inference(usedef, [], [e80])).
fof(e80, plain, (sP25 <=> ((e2 = op(e2, op(e2, e2))) & ~ (op(e2, e2) = op(e2, e2)) & (op(e2, e2) = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f81, plain, (((e3 = op(e3, op(e2, e3))) & ~ (op(e2, e3) = op(e3, e3)) & (op(e2, e3) = op(e3, e2))) | ~ sP26), inference(usedef, [], [e81])).
fof(e81, plain, (sP26 <=> ((e3 = op(e3, op(e2, e3))) & ~ (op(e2, e3) = op(e3, e3)) & (op(e2, e3) = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f82, plain, (((e0 = op(e0, op(e3, e0))) & ~ (op(e0, e0) = op(e3, e0)) & (op(e0, e3) = op(e3, e0))) | ~ sP27), inference(usedef, [], [e82])).
fof(e82, plain, (sP27 <=> ((e0 = op(e0, op(e3, e0))) & ~ (op(e0, e0) = op(e3, e0)) & (op(e0, e3) = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f83, plain, (((e1 = op(e1, op(e3, e1))) & ~ (op(e1, e1) = op(e3, e1)) & (op(e1, e3) = op(e3, e1))) | ~ sP28), inference(usedef, [], [e83])).
fof(e83, plain, (sP28 <=> ((e1 = op(e1, op(e3, e1))) & ~ (op(e1, e1) = op(e3, e1)) & (op(e1, e3) = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f84, plain, (((e2 = op(e2, op(e3, e2))) & ~ (op(e2, e2) = op(e3, e2)) & (op(e2, e3) = op(e3, e2))) | ~ sP29), inference(usedef, [], [e84])).
fof(e84, plain, (sP29 <=> ((e2 = op(e2, op(e3, e2))) & ~ (op(e2, e2) = op(e3, e2)) & (op(e2, e3) = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f85, plain, (((((~ (op(e0, e0) = e3) | ~ (op(e0, e0) = e3)) & (~ (op(e0, e0) = e2) | ~ (op(e0, e0) = e2)) & (~ (op(e0, e0) = e1) | ~ (op(e0, e0) = e1)) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)))) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0))) | ~ sP30), inference(usedef, [], [e85])).
fof(e85, plain, (sP30 <=> ((((~ (op(e0, e0) = e3) | ~ (op(e0, e0) = e3)) & (~ (op(e0, e0) = e2) | ~ (op(e0, e0) = e2)) & (~ (op(e0, e0) = e1) | ~ (op(e0, e0) = e1)) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)))) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f86, plain, (((((~ (e3 = op(e1, e0)) | ~ (e3 = op(e0, e1))) & (~ (e2 = op(e1, e0)) | ~ (e2 = op(e0, e1))) & (~ (e1 = op(e1, e0)) | ~ (e1 = op(e0, e1))) & (~ (e0 = op(e1, e0)) | ~ (e0 = op(e0, e1)))) | ~ (op(e0, e0) = e1)) & (op(e0, e0) = e1)) | ~ sP31), inference(usedef, [], [e86])).
fof(e86, plain, (sP31 <=> ((((~ (e3 = op(e1, e0)) | ~ (e3 = op(e0, e1))) & (~ (e2 = op(e1, e0)) | ~ (e2 = op(e0, e1))) & (~ (e1 = op(e1, e0)) | ~ (e1 = op(e0, e1))) & (~ (e0 = op(e1, e0)) | ~ (e0 = op(e0, e1)))) | ~ (op(e0, e0) = e1)) & (op(e0, e0) = e1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f87, plain, (((((~ (e3 = op(e2, e0)) | ~ (e3 = op(e0, e2))) & (~ (e2 = op(e2, e0)) | ~ (e2 = op(e0, e2))) & (~ (e1 = op(e2, e0)) | ~ (e1 = op(e0, e2))) & (~ (e0 = op(e2, e0)) | ~ (e0 = op(e0, e2)))) | ~ (op(e0, e0) = e2)) & (op(e0, e0) = e2)) | ~ sP32), inference(usedef, [], [e87])).
fof(e87, plain, (sP32 <=> ((((~ (e3 = op(e2, e0)) | ~ (e3 = op(e0, e2))) & (~ (e2 = op(e2, e0)) | ~ (e2 = op(e0, e2))) & (~ (e1 = op(e2, e0)) | ~ (e1 = op(e0, e2))) & (~ (e0 = op(e2, e0)) | ~ (e0 = op(e0, e2)))) | ~ (op(e0, e0) = e2)) & (op(e0, e0) = e2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f88, plain, (((((~ (e3 = op(e3, e0)) | ~ (e3 = op(e0, e3))) & (~ (e2 = op(e3, e0)) | ~ (e2 = op(e0, e3))) & (~ (e1 = op(e3, e0)) | ~ (e1 = op(e0, e3))) & (~ (e0 = op(e3, e0)) | ~ (e0 = op(e0, e3)))) | ~ (op(e0, e0) = e3)) & (op(e0, e0) = e3)) | ~ sP33), inference(usedef, [], [e88])).
fof(e88, plain, (sP33 <=> ((((~ (e3 = op(e3, e0)) | ~ (e3 = op(e0, e3))) & (~ (e2 = op(e3, e0)) | ~ (e2 = op(e0, e3))) & (~ (e1 = op(e3, e0)) | ~ (e1 = op(e0, e3))) & (~ (e0 = op(e3, e0)) | ~ (e0 = op(e0, e3)))) | ~ (op(e0, e0) = e3)) & (op(e0, e0) = e3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f89, plain, (((((~ (e3 = op(e0, e1)) | ~ (e3 = op(e1, e0))) & (~ (e2 = op(e0, e1)) | ~ (e2 = op(e1, e0))) & (~ (e1 = op(e0, e1)) | ~ (e1 = op(e1, e0))) & (~ (e0 = op(e0, e1)) | ~ (e0 = op(e1, e0)))) | ~ (e0 = op(e1, e1))) & (e0 = op(e1, e1))) | ~ sP34), inference(usedef, [], [e89])).
fof(e89, plain, (sP34 <=> ((((~ (e3 = op(e0, e1)) | ~ (e3 = op(e1, e0))) & (~ (e2 = op(e0, e1)) | ~ (e2 = op(e1, e0))) & (~ (e1 = op(e0, e1)) | ~ (e1 = op(e1, e0))) & (~ (e0 = op(e0, e1)) | ~ (e0 = op(e1, e0)))) | ~ (e0 = op(e1, e1))) & (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f90, plain, (((((~ (e3 = op(e1, e1)) | ~ (e3 = op(e1, e1))) & (~ (e2 = op(e1, e1)) | ~ (e2 = op(e1, e1))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e0 = op(e1, e1)))) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1))) | ~ sP35), inference(usedef, [], [e90])).
fof(e90, plain, (sP35 <=> ((((~ (e3 = op(e1, e1)) | ~ (e3 = op(e1, e1))) & (~ (e2 = op(e1, e1)) | ~ (e2 = op(e1, e1))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e0 = op(e1, e1)))) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f91, plain, (((((~ (e3 = op(e2, e1)) | ~ (e3 = op(e1, e2))) & (~ (e2 = op(e2, e1)) | ~ (e2 = op(e1, e2))) & (~ (e1 = op(e2, e1)) | ~ (e1 = op(e1, e2))) & (~ (e0 = op(e2, e1)) | ~ (e0 = op(e1, e2)))) | ~ (e2 = op(e1, e1))) & (e2 = op(e1, e1))) | ~ sP36), inference(usedef, [], [e91])).
fof(e91, plain, (sP36 <=> ((((~ (e3 = op(e2, e1)) | ~ (e3 = op(e1, e2))) & (~ (e2 = op(e2, e1)) | ~ (e2 = op(e1, e2))) & (~ (e1 = op(e2, e1)) | ~ (e1 = op(e1, e2))) & (~ (e0 = op(e2, e1)) | ~ (e0 = op(e1, e2)))) | ~ (e2 = op(e1, e1))) & (e2 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f92, plain, (((((~ (e3 = op(e3, e1)) | ~ (e3 = op(e1, e3))) & (~ (e2 = op(e3, e1)) | ~ (e2 = op(e1, e3))) & (~ (e1 = op(e3, e1)) | ~ (e1 = op(e1, e3))) & (~ (e0 = op(e3, e1)) | ~ (e0 = op(e1, e3)))) | ~ (e3 = op(e1, e1))) & (e3 = op(e1, e1))) | ~ sP37), inference(usedef, [], [e92])).
fof(e92, plain, (sP37 <=> ((((~ (e3 = op(e3, e1)) | ~ (e3 = op(e1, e3))) & (~ (e2 = op(e3, e1)) | ~ (e2 = op(e1, e3))) & (~ (e1 = op(e3, e1)) | ~ (e1 = op(e1, e3))) & (~ (e0 = op(e3, e1)) | ~ (e0 = op(e1, e3)))) | ~ (e3 = op(e1, e1))) & (e3 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f93, plain, (((((~ (e3 = op(e0, e2)) | ~ (e3 = op(e2, e0))) & (~ (e2 = op(e0, e2)) | ~ (e2 = op(e2, e0))) & (~ (e1 = op(e0, e2)) | ~ (e1 = op(e2, e0))) & (~ (e0 = op(e0, e2)) | ~ (e0 = op(e2, e0)))) | ~ (e0 = op(e2, e2))) & (e0 = op(e2, e2))) | ~ sP38), inference(usedef, [], [e93])).
fof(e93, plain, (sP38 <=> ((((~ (e3 = op(e0, e2)) | ~ (e3 = op(e2, e0))) & (~ (e2 = op(e0, e2)) | ~ (e2 = op(e2, e0))) & (~ (e1 = op(e0, e2)) | ~ (e1 = op(e2, e0))) & (~ (e0 = op(e0, e2)) | ~ (e0 = op(e2, e0)))) | ~ (e0 = op(e2, e2))) & (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f94, plain, (((((~ (e3 = op(e1, e2)) | ~ (e3 = op(e2, e1))) & (~ (e2 = op(e1, e2)) | ~ (e2 = op(e2, e1))) & (~ (e1 = op(e1, e2)) | ~ (e1 = op(e2, e1))) & (~ (e0 = op(e1, e2)) | ~ (e0 = op(e2, e1)))) | ~ (e1 = op(e2, e2))) & (e1 = op(e2, e2))) | ~ sP39), inference(usedef, [], [e94])).
fof(e94, plain, (sP39 <=> ((((~ (e3 = op(e1, e2)) | ~ (e3 = op(e2, e1))) & (~ (e2 = op(e1, e2)) | ~ (e2 = op(e2, e1))) & (~ (e1 = op(e1, e2)) | ~ (e1 = op(e2, e1))) & (~ (e0 = op(e1, e2)) | ~ (e0 = op(e2, e1)))) | ~ (e1 = op(e2, e2))) & (e1 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f95, plain, (((((~ (e3 = op(e2, e2)) | ~ (e3 = op(e2, e2))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e1 = op(e2, e2))) & (~ (e0 = op(e2, e2)) | ~ (e0 = op(e2, e2)))) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2))) | ~ sP40), inference(usedef, [], [e95])).
fof(e95, plain, (sP40 <=> ((((~ (e3 = op(e2, e2)) | ~ (e3 = op(e2, e2))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e1 = op(e2, e2))) & (~ (e0 = op(e2, e2)) | ~ (e0 = op(e2, e2)))) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f96, plain, (((((~ (e3 = op(e3, e2)) | ~ (e3 = op(e2, e3))) & (~ (e2 = op(e3, e2)) | ~ (e2 = op(e2, e3))) & (~ (e1 = op(e3, e2)) | ~ (e1 = op(e2, e3))) & (~ (e0 = op(e3, e2)) | ~ (e0 = op(e2, e3)))) | ~ (e3 = op(e2, e2))) & (e3 = op(e2, e2))) | ~ sP41), inference(usedef, [], [e96])).
fof(e96, plain, (sP41 <=> ((((~ (e3 = op(e3, e2)) | ~ (e3 = op(e2, e3))) & (~ (e2 = op(e3, e2)) | ~ (e2 = op(e2, e3))) & (~ (e1 = op(e3, e2)) | ~ (e1 = op(e2, e3))) & (~ (e0 = op(e3, e2)) | ~ (e0 = op(e2, e3)))) | ~ (e3 = op(e2, e2))) & (e3 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f97, plain, (((((~ (e3 = op(e0, e3)) | ~ (e3 = op(e3, e0))) & (~ (e2 = op(e0, e3)) | ~ (e2 = op(e3, e0))) & (~ (e1 = op(e0, e3)) | ~ (e1 = op(e3, e0))) & (~ (e0 = op(e0, e3)) | ~ (e0 = op(e3, e0)))) | ~ (e0 = op(e3, e3))) & (e0 = op(e3, e3))) | ~ sP42), inference(usedef, [], [e97])).
fof(e97, plain, (sP42 <=> ((((~ (e3 = op(e0, e3)) | ~ (e3 = op(e3, e0))) & (~ (e2 = op(e0, e3)) | ~ (e2 = op(e3, e0))) & (~ (e1 = op(e0, e3)) | ~ (e1 = op(e3, e0))) & (~ (e0 = op(e0, e3)) | ~ (e0 = op(e3, e0)))) | ~ (e0 = op(e3, e3))) & (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f98, plain, (((((~ (e3 = op(e1, e3)) | ~ (e3 = op(e3, e1))) & (~ (e2 = op(e1, e3)) | ~ (e2 = op(e3, e1))) & (~ (e1 = op(e1, e3)) | ~ (e1 = op(e3, e1))) & (~ (e0 = op(e1, e3)) | ~ (e0 = op(e3, e1)))) | ~ (e1 = op(e3, e3))) & (e1 = op(e3, e3))) | ~ sP43), inference(usedef, [], [e98])).
fof(e98, plain, (sP43 <=> ((((~ (e3 = op(e1, e3)) | ~ (e3 = op(e3, e1))) & (~ (e2 = op(e1, e3)) | ~ (e2 = op(e3, e1))) & (~ (e1 = op(e1, e3)) | ~ (e1 = op(e3, e1))) & (~ (e0 = op(e1, e3)) | ~ (e0 = op(e3, e1)))) | ~ (e1 = op(e3, e3))) & (e1 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f99, plain, (((((~ (e3 = op(e2, e3)) | ~ (e3 = op(e3, e2))) & (~ (e2 = op(e2, e3)) | ~ (e2 = op(e3, e2))) & (~ (e1 = op(e2, e3)) | ~ (e1 = op(e3, e2))) & (~ (e0 = op(e2, e3)) | ~ (e0 = op(e3, e2)))) | ~ (e2 = op(e3, e3))) & (e2 = op(e3, e3))) | ~ sP44), inference(usedef, [], [e99])).
fof(e99, plain, (sP44 <=> ((((~ (e3 = op(e2, e3)) | ~ (e3 = op(e3, e2))) & (~ (e2 = op(e2, e3)) | ~ (e2 = op(e3, e2))) & (~ (e1 = op(e2, e3)) | ~ (e1 = op(e3, e2))) & (~ (e0 = op(e2, e3)) | ~ (e0 = op(e3, e2)))) | ~ (e2 = op(e3, e3))) & (e2 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f30, plain, ((((((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e2 = op(e3, e3))) & (~ (e1 = op(e3, e3)) | ~ (e1 = op(e3, e3))) & (~ (e0 = op(e3, e3)) | ~ (e0 = op(e3, e3)))) | ~ (e3 = op(e3, e3))) & (e3 = op(e3, e3))) | ((((~ (e3 = op(e2, e3)) | ~ (e3 = op(e3, e2))) & (~ (e2 = op(e2, e3)) | ~ (e2 = op(e3, e2))) & (~ (e1 = op(e2, e3)) | ~ (e1 = op(e3, e2))) & (~ (e0 = op(e2, e3)) | ~ (e0 = op(e3, e2)))) | ~ (e2 = op(e3, e3))) & (e2 = op(e3, e3))) | ((((~ (e3 = op(e1, e3)) | ~ (e3 = op(e3, e1))) & (~ (e2 = op(e1, e3)) | ~ (e2 = op(e3, e1))) & (~ (e1 = op(e1, e3)) | ~ (e1 = op(e3, e1))) & (~ (e0 = op(e1, e3)) | ~ (e0 = op(e3, e1)))) | ~ (e1 = op(e3, e3))) & (e1 = op(e3, e3))) | ((((~ (e3 = op(e0, e3)) | ~ (e3 = op(e3, e0))) & (~ (e2 = op(e0, e3)) | ~ (e2 = op(e3, e0))) & (~ (e1 = op(e0, e3)) | ~ (e1 = op(e3, e0))) & (~ (e0 = op(e0, e3)) | ~ (e0 = op(e3, e0)))) | ~ (e0 = op(e3, e3))) & (e0 = op(e3, e3))) | ((((~ (e3 = op(e3, e2)) | ~ (e3 = op(e2, e3))) & (~ (e2 = op(e3, e2)) | ~ (e2 = op(e2, e3))) & (~ (e1 = op(e3, e2)) | ~ (e1 = op(e2, e3))) & (~ (e0 = op(e3, e2)) | ~ (e0 = op(e2, e3)))) | ~ (e3 = op(e2, e2))) & (e3 = op(e2, e2))) | ((((~ (e3 = op(e2, e2)) | ~ (e3 = op(e2, e2))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e1 = op(e2, e2))) & (~ (e0 = op(e2, e2)) | ~ (e0 = op(e2, e2)))) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2))) | ((((~ (e3 = op(e1, e2)) | ~ (e3 = op(e2, e1))) & (~ (e2 = op(e1, e2)) | ~ (e2 = op(e2, e1))) & (~ (e1 = op(e1, e2)) | ~ (e1 = op(e2, e1))) & (~ (e0 = op(e1, e2)) | ~ (e0 = op(e2, e1)))) | ~ (e1 = op(e2, e2))) & (e1 = op(e2, e2))) | ((((~ (e3 = op(e0, e2)) | ~ (e3 = op(e2, e0))) & (~ (e2 = op(e0, e2)) | ~ (e2 = op(e2, e0))) & (~ (e1 = op(e0, e2)) | ~ (e1 = op(e2, e0))) & (~ (e0 = op(e0, e2)) | ~ (e0 = op(e2, e0)))) | ~ (e0 = op(e2, e2))) & (e0 = op(e2, e2))) | ((((~ (e3 = op(e3, e1)) | ~ (e3 = op(e1, e3))) & (~ (e2 = op(e3, e1)) | ~ (e2 = op(e1, e3))) & (~ (e1 = op(e3, e1)) | ~ (e1 = op(e1, e3))) & (~ (e0 = op(e3, e1)) | ~ (e0 = op(e1, e3)))) | ~ (e3 = op(e1, e1))) & (e3 = op(e1, e1))) | ((((~ (e3 = op(e2, e1)) | ~ (e3 = op(e1, e2))) & (~ (e2 = op(e2, e1)) | ~ (e2 = op(e1, e2))) & (~ (e1 = op(e2, e1)) | ~ (e1 = op(e1, e2))) & (~ (e0 = op(e2, e1)) | ~ (e0 = op(e1, e2)))) | ~ (e2 = op(e1, e1))) & (e2 = op(e1, e1))) | ((((~ (e3 = op(e1, e1)) | ~ (e3 = op(e1, e1))) & (~ (e2 = op(e1, e1)) | ~ (e2 = op(e1, e1))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e0 = op(e1, e1)))) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1))) | ((((~ (e3 = op(e0, e1)) | ~ (e3 = op(e1, e0))) & (~ (e2 = op(e0, e1)) | ~ (e2 = op(e1, e0))) & (~ (e1 = op(e0, e1)) | ~ (e1 = op(e1, e0))) & (~ (e0 = op(e0, e1)) | ~ (e0 = op(e1, e0)))) | ~ (e0 = op(e1, e1))) & (e0 = op(e1, e1))) | ((((~ (e3 = op(e3, e0)) | ~ (e3 = op(e0, e3))) & (~ (e2 = op(e3, e0)) | ~ (e2 = op(e0, e3))) & (~ (e1 = op(e3, e0)) | ~ (e1 = op(e0, e3))) & (~ (e0 = op(e3, e0)) | ~ (e0 = op(e0, e3)))) | ~ (op(e0, e0) = e3)) & (op(e0, e0) = e3)) | ((((~ (e3 = op(e2, e0)) | ~ (e3 = op(e0, e2))) & (~ (e2 = op(e2, e0)) | ~ (e2 = op(e0, e2))) & (~ (e1 = op(e2, e0)) | ~ (e1 = op(e0, e2))) & (~ (e0 = op(e2, e0)) | ~ (e0 = op(e0, e2)))) | ~ (op(e0, e0) = e2)) & (op(e0, e0) = e2)) | ((((~ (e3 = op(e1, e0)) | ~ (e3 = op(e0, e1))) & (~ (e2 = op(e1, e0)) | ~ (e2 = op(e0, e1))) & (~ (e1 = op(e1, e0)) | ~ (e1 = op(e0, e1))) & (~ (e0 = op(e1, e0)) | ~ (e0 = op(e0, e1)))) | ~ (op(e0, e0) = e1)) & (op(e0, e0) = e1)) | ((((~ (op(e0, e0) = e3) | ~ (op(e0, e0) = e3)) & (~ (op(e0, e0) = e2) | ~ (op(e0, e0) = e2)) & (~ (op(e0, e0) = e1) | ~ (op(e0, e0) = e1)) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)))) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)))) & (((e3 = op(e3, op(e3, e3))) & ~ (op(e3, e3) = op(e3, e3)) & (op(e3, e3) = op(e3, e3))) | ((e2 = op(e2, op(e3, e2))) & ~ (op(e2, e2) = op(e3, e2)) & (op(e2, e3) = op(e3, e2))) | ((e1 = op(e1, op(e3, e1))) & ~ (op(e1, e1) = op(e3, e1)) & (op(e1, e3) = op(e3, e1))) | ((e0 = op(e0, op(e3, e0))) & ~ (op(e0, e0) = op(e3, e0)) & (op(e0, e3) = op(e3, e0))) | ((e3 = op(e3, op(e2, e3))) & ~ (op(e2, e3) = op(e3, e3)) & (op(e2, e3) = op(e3, e2))) | ((e2 = op(e2, op(e2, e2))) & ~ (op(e2, e2) = op(e2, e2)) & (op(e2, e2) = op(e2, e2))) | ((e1 = op(e1, op(e2, e1))) & ~ (op(e1, e1) = op(e2, e1)) & (op(e1, e2) = op(e2, e1))) | ((e0 = op(e0, op(e2, e0))) & ~ (op(e0, e0) = op(e2, e0)) & (op(e0, e2) = op(e2, e0))) | ((e3 = op(e3, op(e1, e3))) & ~ (op(e1, e3) = op(e3, e3)) & (op(e1, e3) = op(e3, e1))) | ((e2 = op(e2, op(e1, e2))) & ~ (op(e1, e2) = op(e2, e2)) & (op(e1, e2) = op(e2, e1))) | ((e1 = op(e1, op(e1, e1))) & ~ (op(e1, e1) = op(e1, e1)) & (op(e1, e1) = op(e1, e1))) | ((e0 = op(e0, op(e1, e0))) & ~ (op(e0, e0) = op(e1, e0)) & (op(e0, e1) = op(e1, e0))) | ((e3 = op(e3, op(e0, e3))) & ~ (op(e0, e3) = op(e3, e3)) & (op(e0, e3) = op(e3, e0))) | ((e2 = op(e2, op(e0, e2))) & ~ (op(e0, e2) = op(e2, e2)) & (op(e0, e2) = op(e2, e0))) | ((e1 = op(e1, op(e0, e1))) & ~ (op(e0, e1) = op(e1, e1)) & (op(e0, e1) = op(e1, e0))) | ((e0 = op(e0, op(e0, e0))) & ~ (op(e0, e0) = op(e0, e0)) & (op(e0, e0) = op(e0, e0)))) & ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | (~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | (~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, (((~ ((((e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e2 = op(e3, e3)) & (e2 = op(e3, e3))) | ((e1 = op(e3, e3)) & (e1 = op(e3, e3))) | ((e0 = op(e3, e3)) & (e0 = op(e3, e3)))) & (e3 = op(e3, e3))) & (e3 = op(e3, e3))) | (~ ((((e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((e2 = op(e2, e3)) & (e2 = op(e3, e2))) | ((e1 = op(e2, e3)) & (e1 = op(e3, e2))) | ((e0 = op(e2, e3)) & (e0 = op(e3, e2)))) & (e2 = op(e3, e3))) & (e2 = op(e3, e3))) | (~ ((((e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((e2 = op(e1, e3)) & (e2 = op(e3, e1))) | ((e1 = op(e1, e3)) & (e1 = op(e3, e1))) | ((e0 = op(e1, e3)) & (e0 = op(e3, e1)))) & (e1 = op(e3, e3))) & (e1 = op(e3, e3))) | (~ ((((e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((e2 = op(e0, e3)) & (e2 = op(e3, e0))) | ((e1 = op(e0, e3)) & (e1 = op(e3, e0))) | ((e0 = op(e0, e3)) & (e0 = op(e3, e0)))) & (e0 = op(e3, e3))) & (e0 = op(e3, e3))) | (~ ((((e3 = op(e3, e2)) & (e3 = op(e2, e3))) | ((e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((e1 = op(e3, e2)) & (e1 = op(e2, e3))) | ((e0 = op(e3, e2)) & (e0 = op(e2, e3)))) & (e3 = op(e2, e2))) & (e3 = op(e2, e2))) | (~ ((((e3 = op(e2, e2)) & (e3 = op(e2, e2))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e1 = op(e2, e2)) & (e1 = op(e2, e2))) | ((e0 = op(e2, e2)) & (e0 = op(e2, e2)))) & (e2 = op(e2, e2))) & (e2 = op(e2, e2))) | (~ ((((e3 = op(e1, e2)) & (e3 = op(e2, e1))) | ((e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((e1 = op(e1, e2)) & (e1 = op(e2, e1))) | ((e0 = op(e1, e2)) & (e0 = op(e2, e1)))) & (e1 = op(e2, e2))) & (e1 = op(e2, e2))) | (~ ((((e3 = op(e0, e2)) & (e3 = op(e2, e0))) | ((e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((e1 = op(e0, e2)) & (e1 = op(e2, e0))) | ((e0 = op(e0, e2)) & (e0 = op(e2, e0)))) & (e0 = op(e2, e2))) & (e0 = op(e2, e2))) | (~ ((((e3 = op(e3, e1)) & (e3 = op(e1, e3))) | ((e2 = op(e3, e1)) & (e2 = op(e1, e3))) | ((e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((e0 = op(e3, e1)) & (e0 = op(e1, e3)))) & (e3 = op(e1, e1))) & (e3 = op(e1, e1))) | (~ ((((e3 = op(e2, e1)) & (e3 = op(e1, e2))) | ((e2 = op(e2, e1)) & (e2 = op(e1, e2))) | ((e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((e0 = op(e2, e1)) & (e0 = op(e1, e2)))) & (e2 = op(e1, e1))) & (e2 = op(e1, e1))) | (~ ((((e3 = op(e1, e1)) & (e3 = op(e1, e1))) | ((e2 = op(e1, e1)) & (e2 = op(e1, e1))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e0 = op(e1, e1)) & (e0 = op(e1, e1)))) & (e1 = op(e1, e1))) & (e1 = op(e1, e1))) | (~ ((((e3 = op(e0, e1)) & (e3 = op(e1, e0))) | ((e2 = op(e0, e1)) & (e2 = op(e1, e0))) | ((e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((e0 = op(e0, e1)) & (e0 = op(e1, e0)))) & (e0 = op(e1, e1))) & (e0 = op(e1, e1))) | (~ ((((e3 = op(e3, e0)) & (e3 = op(e0, e3))) | ((e2 = op(e3, e0)) & (e2 = op(e0, e3))) | ((e1 = op(e3, e0)) & (e1 = op(e0, e3))) | ((e0 = op(e3, e0)) & (e0 = op(e0, e3)))) & (op(e0, e0) = e3)) & (op(e0, e0) = e3)) | (~ ((((e3 = op(e2, e0)) & (e3 = op(e0, e2))) | ((e2 = op(e2, e0)) & (e2 = op(e0, e2))) | ((e1 = op(e2, e0)) & (e1 = op(e0, e2))) | ((e0 = op(e2, e0)) & (e0 = op(e0, e2)))) & (op(e0, e0) = e2)) & (op(e0, e0) = e2)) | (~ ((((e3 = op(e1, e0)) & (e3 = op(e0, e1))) | ((e2 = op(e1, e0)) & (e2 = op(e0, e1))) | ((e1 = op(e1, e0)) & (e1 = op(e0, e1))) | ((e0 = op(e1, e0)) & (e0 = op(e0, e1)))) & (op(e0, e0) = e1)) & (op(e0, e0) = e1)) | (~ ((((op(e0, e0) = e3) & (op(e0, e0) = e3)) | ((op(e0, e0) = e2) & (op(e0, e0) = e2)) | ((op(e0, e0) = e1) & (op(e0, e0) = e1)) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & (e0 = op(e0, e0))) & (e0 = op(e0, e0)))) & (((e3 = op(e3, op(e3, e3))) & ~ (op(e3, e3) = op(e3, e3)) & (op(e3, e3) = op(e3, e3))) | ((e2 = op(e2, op(e3, e2))) & ~ (op(e2, e2) = op(e3, e2)) & (op(e2, e3) = op(e3, e2))) | ((e1 = op(e1, op(e3, e1))) & ~ (op(e1, e1) = op(e3, e1)) & (op(e1, e3) = op(e3, e1))) | ((e0 = op(e0, op(e3, e0))) & ~ (op(e0, e0) = op(e3, e0)) & (op(e0, e3) = op(e3, e0))) | ((e3 = op(e3, op(e2, e3))) & ~ (op(e2, e3) = op(e3, e3)) & (op(e2, e3) = op(e3, e2))) | ((e2 = op(e2, op(e2, e2))) & ~ (op(e2, e2) = op(e2, e2)) & (op(e2, e2) = op(e2, e2))) | ((e1 = op(e1, op(e2, e1))) & ~ (op(e1, e1) = op(e2, e1)) & (op(e1, e2) = op(e2, e1))) | ((e0 = op(e0, op(e2, e0))) & ~ (op(e0, e0) = op(e2, e0)) & (op(e0, e2) = op(e2, e0))) | ((e3 = op(e3, op(e1, e3))) & ~ (op(e1, e3) = op(e3, e3)) & (op(e1, e3) = op(e3, e1))) | ((e2 = op(e2, op(e1, e2))) & ~ (op(e1, e2) = op(e2, e2)) & (op(e1, e2) = op(e2, e1))) | ((e1 = op(e1, op(e1, e1))) & ~ (op(e1, e1) = op(e1, e1)) & (op(e1, e1) = op(e1, e1))) | ((e0 = op(e0, op(e1, e0))) & ~ (op(e0, e0) = op(e1, e0)) & (op(e0, e1) = op(e1, e0))) | ((e3 = op(e3, op(e0, e3))) & ~ (op(e0, e3) = op(e3, e3)) & (op(e0, e3) = op(e3, e0))) | ((e2 = op(e2, op(e0, e2))) & ~ (op(e0, e2) = op(e2, e2)) & (op(e0, e2) = op(e2, e0))) | ((e1 = op(e1, op(e0, e1))) & ~ (op(e0, e1) = op(e1, e1)) & (op(e0, e1) = op(e1, e0))) | ((e0 = op(e0, op(e0, e0))) & ~ (op(e0, e0) = op(e0, e0)) & (op(e0, e0) = op(e0, e0)))) & ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | (~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | (~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | (~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | (~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax5)).
fof(f1293, plain, (spl45_130 | spl45_129 | spl45_128 | spl45_127 | spl45_126 | spl45_125 | spl45_124 | spl45_123 | spl45_122 | spl45_121 | spl45_120 | spl45_119 | spl45_118 | spl45_117 | spl45_116 | ~ spl45_49), inference(avatar_split_clause, [], [f460, f721, f1130, f1140, f1150, f1160, f1170, f1180, f1190, f1200, f1210, f1220, f1230, f1240, f1250, f1260, f1270])).
fof(f460, plain, (~ (e0 = op(e0, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f100])).
fof(f1292, plain, (spl45_130 | spl45_129 | spl45_128 | spl45_127 | spl45_126 | spl45_125 | spl45_124 | spl45_123 | spl45_122 | spl45_121 | spl45_120 | spl45_119 | spl45_118 | spl45_117 | spl45_116 | ~ spl45_34), inference(avatar_split_clause, [], [f461, f657, f1130, f1140, f1150, f1160, f1170, f1180, f1190, f1200, f1210, f1220, f1230, f1240, f1250, f1260, f1270])).
fof(f461, plain, (~ (e1 = op(e1, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f100])).
fof(f1291, plain, (spl45_130 | spl45_129 | spl45_128 | spl45_127 | spl45_126 | spl45_125 | spl45_124 | spl45_123 | spl45_122 | spl45_121 | spl45_120 | spl45_119 | spl45_118 | spl45_117 | spl45_116 | ~ spl45_19), inference(avatar_split_clause, [], [f462, f593, f1130, f1140, f1150, f1160, f1170, f1180, f1190, f1200, f1210, f1220, f1230, f1240, f1250, f1260, f1270])).
fof(f462, plain, (~ (e2 = op(e2, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f100])).
fof(f1290, plain, (spl45_130 | spl45_129 | spl45_128 | spl45_127 | spl45_126 | spl45_125 | spl45_124 | spl45_123 | spl45_122 | spl45_121 | spl45_120 | spl45_119 | spl45_118 | spl45_117 | spl45_116 | ~ spl45_4), inference(avatar_split_clause, [], [f463, f529, f1130, f1140, f1150, f1160, f1170, f1180, f1190, f1200, f1210, f1220, f1230, f1240, f1250, f1260, f1270])).
fof(f463, plain, (~ (e3 = op(e3, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f100])).
fof(f1289, plain, (spl45_114 | spl45_112 | spl45_110 | spl45_108 | spl45_105 | spl45_103 | spl45_101 | spl45_99 | spl45_96 | spl45_93 | spl45_91 | spl45_89 | spl45_86 | spl45_83 | spl45_80), inference(avatar_split_clause, [], [f496, f956, f970, f984, f998, f1008, f1018, f1032, f1046, f1056, f1066, f1076, f1090, f1100, f1110, f1120])).
fof(f1120, plain, (spl45_114 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl45_114])])).
fof(f1110, plain, (spl45_112 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl45_112])])).
fof(f1100, plain, (spl45_110 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl45_110])])).
fof(f1090, plain, (spl45_108 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl45_108])])).
fof(f1076, plain, (spl45_105 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl45_105])])).
fof(f1066, plain, (spl45_103 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl45_103])])).
fof(f1056, plain, (spl45_101 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl45_101])])).
fof(f1046, plain, (spl45_99 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl45_99])])).
fof(f1032, plain, (spl45_96 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl45_96])])).
fof(f1018, plain, (spl45_93 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl45_93])])).
fof(f1008, plain, (spl45_91 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl45_91])])).
fof(f998, plain, (spl45_89 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl45_89])])).
fof(f984, plain, (spl45_86 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl45_86])])).
fof(f970, plain, (spl45_83 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl45_83])])).
fof(f956, plain, (spl45_80 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl45_80])])).
fof(f496, plain, (sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(trivial_inequality_removal, [], [f465])).
fof(f465, plain, (~ (op(e3, e3) = op(e3, e3)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f100])).
fof(f1283, plain, (spl45_79 | spl45_78 | spl45_77 | spl45_76 | spl45_75 | spl45_74 | spl45_73 | spl45_72 | spl45_71 | spl45_70 | spl45_69 | spl45_68 | spl45_67 | spl45_66 | spl45_65 | spl45_4), inference(avatar_split_clause, [], [f467, f529, f821, f830, f839, f848, f857, f866, f875, f884, f893, f902, f911, f920, f929, f938, f947])).
fof(f947, plain, (spl45_79 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl45_79])])).
fof(f938, plain, (spl45_78 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl45_78])])).
fof(f929, plain, (spl45_77 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl45_77])])).
fof(f920, plain, (spl45_76 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl45_76])])).
fof(f911, plain, (spl45_75 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl45_75])])).
fof(f902, plain, (spl45_74 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl45_74])])).
fof(f893, plain, (spl45_73 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl45_73])])).
fof(f884, plain, (spl45_72 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl45_72])])).
fof(f875, plain, (spl45_71 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl45_71])])).
fof(f866, plain, (spl45_70 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl45_70])])).
fof(f857, plain, (spl45_69 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl45_69])])).
fof(f848, plain, (spl45_68 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl45_68])])).
fof(f839, plain, (spl45_67 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl45_67])])).
fof(f830, plain, (spl45_66 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl45_66])])).
fof(f821, plain, (spl45_65 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl45_65])])).
fof(f467, plain, ((e3 = op(e3, e3)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(cnf_transformation, [], [f100])).
fof(f1279, plain, (spl45_79 | spl45_78 | spl45_77 | spl45_76 | spl45_75 | spl45_74 | spl45_73 | spl45_72 | spl45_71 | spl45_70 | spl45_69 | spl45_68 | spl45_67 | spl45_66 | spl45_65 | ~ spl45_4), inference(avatar_split_clause, [], [f500, f529, f821, f830, f839, f848, f857, f866, f875, f884, f893, f902, f911, f920, f929, f938, f947])).
fof(f500, plain, (~ (e3 = op(e3, e3)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(duplicate_literal_removal, [], [f471])).
fof(f471, plain, (~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(cnf_transformation, [], [f100])).
fof(f1278, plain, (~ spl45_130 | spl45_61), inference(avatar_split_clause, [], [f452, f772, f1270])).
fof(f452, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f55])).
fof(f1276, plain, (~ spl45_130 | ~ spl45_61), inference(avatar_split_clause, [], [f454, f772, f1270])).
fof(f454, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f145])).
fof(f1268, plain, (~ spl45_129 | spl45_57), inference(avatar_split_clause, [], [f446, f755, f1260])).
fof(f446, plain, ((e0 = op(e0, e1)) | ~ sP1), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP1), inference(nnf_transformation, [], [f56])).
fof(f1267, plain, (~ spl45_129 | spl45_45), inference(avatar_split_clause, [], [f447, f704, f1260])).
fof(f447, plain, ((e0 = op(e1, e0)) | ~ sP1), inference(cnf_transformation, [], [f144])).
fof(f1266, plain, (~ spl45_129 | ~ spl45_61), inference(avatar_split_clause, [], [f448, f772, f1260])).
fof(f448, plain, (~ (e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f144])).
fof(f1265, plain, (~ spl45_129 | ~ spl45_46), inference(avatar_split_clause, [], [f449, f708, f1260])).
fof(f449, plain, (~ (e1 = op(e1, e0)) | ~ sP1), inference(cnf_transformation, [], [f144])).
fof(f1264, plain, (~ spl45_129 | ~ spl45_31), inference(avatar_split_clause, [], [f450, f644, f1260])).
fof(f450, plain, (~ (e2 = op(e2, e0)) | ~ sP1), inference(cnf_transformation, [], [f144])).
fof(f1263, plain, (~ spl45_129 | ~ spl45_16), inference(avatar_split_clause, [], [f451, f580, f1260])).
fof(f451, plain, (~ (e3 = op(e3, e0)) | ~ sP1), inference(cnf_transformation, [], [f144])).
fof(f1258, plain, (~ spl45_128 | spl45_53), inference(avatar_split_clause, [], [f440, f738, f1250])).
fof(f440, plain, ((e0 = op(e0, e2)) | ~ sP2), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP2), inference(nnf_transformation, [], [f57])).
fof(f1257, plain, (~ spl45_128 | spl45_29), inference(avatar_split_clause, [], [f441, f636, f1250])).
fof(f441, plain, ((e0 = op(e2, e0)) | ~ sP2), inference(cnf_transformation, [], [f143])).
fof(f1256, plain, (~ spl45_128 | ~ spl45_61), inference(avatar_split_clause, [], [f442, f772, f1250])).
fof(f442, plain, (~ (e0 = op(e0, e0)) | ~ sP2), inference(cnf_transformation, [], [f143])).
fof(f1255, plain, (~ spl45_128 | ~ spl45_46), inference(avatar_split_clause, [], [f443, f708, f1250])).
fof(f443, plain, (~ (e1 = op(e1, e0)) | ~ sP2), inference(cnf_transformation, [], [f143])).
fof(f1254, plain, (~ spl45_128 | ~ spl45_31), inference(avatar_split_clause, [], [f444, f644, f1250])).
fof(f444, plain, (~ (e2 = op(e2, e0)) | ~ sP2), inference(cnf_transformation, [], [f143])).
fof(f1253, plain, (~ spl45_128 | ~ spl45_16), inference(avatar_split_clause, [], [f445, f580, f1250])).
fof(f445, plain, (~ (e3 = op(e3, e0)) | ~ sP2), inference(cnf_transformation, [], [f143])).
fof(f1248, plain, (~ spl45_127 | spl45_49), inference(avatar_split_clause, [], [f434, f721, f1240])).
fof(f434, plain, ((e0 = op(e0, e3)) | ~ sP3), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e3 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP3), inference(nnf_transformation, [], [f58])).
fof(f1247, plain, (~ spl45_127 | spl45_13), inference(avatar_split_clause, [], [f435, f568, f1240])).
fof(f435, plain, ((e0 = op(e3, e0)) | ~ sP3), inference(cnf_transformation, [], [f142])).
fof(f1246, plain, (~ spl45_127 | ~ spl45_61), inference(avatar_split_clause, [], [f436, f772, f1240])).
fof(f436, plain, (~ (e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f142])).
fof(f1245, plain, (~ spl45_127 | ~ spl45_46), inference(avatar_split_clause, [], [f437, f708, f1240])).
fof(f437, plain, (~ (e1 = op(e1, e0)) | ~ sP3), inference(cnf_transformation, [], [f142])).
fof(f1244, plain, (~ spl45_127 | ~ spl45_31), inference(avatar_split_clause, [], [f438, f644, f1240])).
fof(f438, plain, (~ (e2 = op(e2, e0)) | ~ sP3), inference(cnf_transformation, [], [f142])).
fof(f1243, plain, (~ spl45_127 | ~ spl45_16), inference(avatar_split_clause, [], [f439, f580, f1240])).
fof(f439, plain, (~ (e3 = op(e3, e0)) | ~ sP3), inference(cnf_transformation, [], [f142])).
fof(f1238, plain, (~ spl45_126 | spl45_46), inference(avatar_split_clause, [], [f428, f708, f1230])).
fof(f428, plain, ((e1 = op(e1, e0)) | ~ sP4), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP4), inference(nnf_transformation, [], [f59])).
fof(f1237, plain, (~ spl45_126 | spl45_58), inference(avatar_split_clause, [], [f429, f759, f1230])).
fof(f429, plain, ((e1 = op(e0, e1)) | ~ sP4), inference(cnf_transformation, [], [f141])).
fof(f1236, plain, (~ spl45_126 | ~ spl45_57), inference(avatar_split_clause, [], [f430, f755, f1230])).
fof(f430, plain, (~ (e0 = op(e0, e1)) | ~ sP4), inference(cnf_transformation, [], [f141])).
fof(f1234, plain, (~ spl45_126 | ~ spl45_27), inference(avatar_split_clause, [], [f432, f627, f1230])).
fof(f432, plain, (~ (e2 = op(e2, e1)) | ~ sP4), inference(cnf_transformation, [], [f141])).
fof(f1233, plain, (~ spl45_126 | ~ spl45_12), inference(avatar_split_clause, [], [f433, f563, f1230])).
fof(f433, plain, (~ (e3 = op(e3, e1)) | ~ sP4), inference(cnf_transformation, [], [f141])).
fof(f1228, plain, (~ spl45_125 | spl45_42), inference(avatar_split_clause, [], [f422, f691, f1220])).
fof(f422, plain, ((e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP5), inference(nnf_transformation, [], [f60])).
fof(f1225, plain, (~ spl45_125 | ~ spl45_42), inference(avatar_split_clause, [], [f425, f691, f1220])).
fof(f425, plain, (~ (e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f140])).
fof(f1218, plain, (~ spl45_124 | spl45_38), inference(avatar_split_clause, [], [f416, f674, f1210])).
fof(f416, plain, ((e1 = op(e1, e2)) | ~ sP6), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP6), inference(nnf_transformation, [], [f61])).
fof(f1217, plain, (~ spl45_124 | spl45_26), inference(avatar_split_clause, [], [f417, f623, f1210])).
fof(f417, plain, ((e1 = op(e2, e1)) | ~ sP6), inference(cnf_transformation, [], [f139])).
fof(f1216, plain, (~ spl45_124 | ~ spl45_57), inference(avatar_split_clause, [], [f418, f755, f1210])).
fof(f418, plain, (~ (e0 = op(e0, e1)) | ~ sP6), inference(cnf_transformation, [], [f139])).
fof(f1215, plain, (~ spl45_124 | ~ spl45_42), inference(avatar_split_clause, [], [f419, f691, f1210])).
fof(f419, plain, (~ (e1 = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f139])).
fof(f1214, plain, (~ spl45_124 | ~ spl45_27), inference(avatar_split_clause, [], [f420, f627, f1210])).
fof(f420, plain, (~ (e2 = op(e2, e1)) | ~ sP6), inference(cnf_transformation, [], [f139])).
fof(f1213, plain, (~ spl45_124 | ~ spl45_12), inference(avatar_split_clause, [], [f421, f563, f1210])).
fof(f421, plain, (~ (e3 = op(e3, e1)) | ~ sP6), inference(cnf_transformation, [], [f139])).
fof(f1208, plain, (~ spl45_123 | spl45_34), inference(avatar_split_clause, [], [f410, f657, f1200])).
fof(f410, plain, ((e1 = op(e1, e3)) | ~ sP7), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e3 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e1)) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP7), inference(nnf_transformation, [], [f62])).
fof(f1207, plain, (~ spl45_123 | spl45_10), inference(avatar_split_clause, [], [f411, f555, f1200])).
fof(f411, plain, ((e1 = op(e3, e1)) | ~ sP7), inference(cnf_transformation, [], [f138])).
fof(f1206, plain, (~ spl45_123 | ~ spl45_57), inference(avatar_split_clause, [], [f412, f755, f1200])).
fof(f412, plain, (~ (e0 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f138])).
fof(f1205, plain, (~ spl45_123 | ~ spl45_42), inference(avatar_split_clause, [], [f413, f691, f1200])).
fof(f413, plain, (~ (e1 = op(e1, e1)) | ~ sP7), inference(cnf_transformation, [], [f138])).
fof(f1204, plain, (~ spl45_123 | ~ spl45_27), inference(avatar_split_clause, [], [f414, f627, f1200])).
fof(f414, plain, (~ (e2 = op(e2, e1)) | ~ sP7), inference(cnf_transformation, [], [f138])).
fof(f1203, plain, (~ spl45_123 | ~ spl45_12), inference(avatar_split_clause, [], [f415, f563, f1200])).
fof(f415, plain, (~ (e3 = op(e3, e1)) | ~ sP7), inference(cnf_transformation, [], [f138])).
fof(f1198, plain, (~ spl45_122 | spl45_31), inference(avatar_split_clause, [], [f404, f644, f1190])).
fof(f404, plain, ((e2 = op(e2, e0)) | ~ sP8), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP8), inference(nnf_transformation, [], [f63])).
fof(f1197, plain, (~ spl45_122 | spl45_55), inference(avatar_split_clause, [], [f405, f746, f1190])).
fof(f405, plain, ((e2 = op(e0, e2)) | ~ sP8), inference(cnf_transformation, [], [f137])).
fof(f1196, plain, (~ spl45_122 | ~ spl45_53), inference(avatar_split_clause, [], [f406, f738, f1190])).
fof(f406, plain, (~ (e0 = op(e0, e2)) | ~ sP8), inference(cnf_transformation, [], [f137])).
fof(f1195, plain, (~ spl45_122 | ~ spl45_38), inference(avatar_split_clause, [], [f407, f674, f1190])).
fof(f407, plain, (~ (e1 = op(e1, e2)) | ~ sP8), inference(cnf_transformation, [], [f137])).
fof(f1194, plain, (~ spl45_122 | ~ spl45_23), inference(avatar_split_clause, [], [f408, f610, f1190])).
fof(f408, plain, (~ (e2 = op(e2, e2)) | ~ sP8), inference(cnf_transformation, [], [f137])).
fof(f1193, plain, (~ spl45_122 | ~ spl45_8), inference(avatar_split_clause, [], [f409, f546, f1190])).
fof(f409, plain, (~ (e3 = op(e3, e2)) | ~ sP8), inference(cnf_transformation, [], [f137])).
fof(f1188, plain, (~ spl45_121 | spl45_27), inference(avatar_split_clause, [], [f398, f627, f1180])).
fof(f398, plain, ((e2 = op(e2, e1)) | ~ sP9), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP9), inference(nnf_transformation, [], [f64])).
fof(f1187, plain, (~ spl45_121 | spl45_39), inference(avatar_split_clause, [], [f399, f678, f1180])).
fof(f399, plain, ((e2 = op(e1, e2)) | ~ sP9), inference(cnf_transformation, [], [f136])).
fof(f1186, plain, (~ spl45_121 | ~ spl45_53), inference(avatar_split_clause, [], [f400, f738, f1180])).
fof(f400, plain, (~ (e0 = op(e0, e2)) | ~ sP9), inference(cnf_transformation, [], [f136])).
fof(f1185, plain, (~ spl45_121 | ~ spl45_38), inference(avatar_split_clause, [], [f401, f674, f1180])).
fof(f401, plain, (~ (e1 = op(e1, e2)) | ~ sP9), inference(cnf_transformation, [], [f136])).
fof(f1184, plain, (~ spl45_121 | ~ spl45_23), inference(avatar_split_clause, [], [f402, f610, f1180])).
fof(f402, plain, (~ (e2 = op(e2, e2)) | ~ sP9), inference(cnf_transformation, [], [f136])).
fof(f1183, plain, (~ spl45_121 | ~ spl45_8), inference(avatar_split_clause, [], [f403, f546, f1180])).
fof(f403, plain, (~ (e3 = op(e3, e2)) | ~ sP9), inference(cnf_transformation, [], [f136])).
fof(f1178, plain, (~ spl45_120 | spl45_23), inference(avatar_split_clause, [], [f392, f610, f1170])).
fof(f392, plain, ((e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP10), inference(nnf_transformation, [], [f65])).
fof(f1174, plain, (~ spl45_120 | ~ spl45_23), inference(avatar_split_clause, [], [f396, f610, f1170])).
fof(f396, plain, (~ (e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f135])).
fof(f1168, plain, (~ spl45_119 | spl45_19), inference(avatar_split_clause, [], [f386, f593, f1160])).
fof(f386, plain, ((e2 = op(e2, e3)) | ~ sP11), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e3 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e0, e2)) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP11), inference(nnf_transformation, [], [f66])).
fof(f1167, plain, (~ spl45_119 | spl45_7), inference(avatar_split_clause, [], [f387, f542, f1160])).
fof(f387, plain, ((e2 = op(e3, e2)) | ~ sP11), inference(cnf_transformation, [], [f134])).
fof(f1166, plain, (~ spl45_119 | ~ spl45_53), inference(avatar_split_clause, [], [f388, f738, f1160])).
fof(f388, plain, (~ (e0 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f134])).
fof(f1165, plain, (~ spl45_119 | ~ spl45_38), inference(avatar_split_clause, [], [f389, f674, f1160])).
fof(f389, plain, (~ (e1 = op(e1, e2)) | ~ sP11), inference(cnf_transformation, [], [f134])).
fof(f1164, plain, (~ spl45_119 | ~ spl45_23), inference(avatar_split_clause, [], [f390, f610, f1160])).
fof(f390, plain, (~ (e2 = op(e2, e2)) | ~ sP11), inference(cnf_transformation, [], [f134])).
fof(f1163, plain, (~ spl45_119 | ~ spl45_8), inference(avatar_split_clause, [], [f391, f546, f1160])).
fof(f391, plain, (~ (e3 = op(e3, e2)) | ~ sP11), inference(cnf_transformation, [], [f134])).
fof(f1158, plain, (~ spl45_118 | spl45_16), inference(avatar_split_clause, [], [f380, f580, f1150])).
fof(f380, plain, ((e3 = op(e3, e0)) | ~ sP12), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP12), inference(nnf_transformation, [], [f67])).
fof(f1157, plain, (~ spl45_118 | spl45_52), inference(avatar_split_clause, [], [f381, f733, f1150])).
fof(f381, plain, ((e3 = op(e0, e3)) | ~ sP12), inference(cnf_transformation, [], [f133])).
fof(f1156, plain, (~ spl45_118 | ~ spl45_49), inference(avatar_split_clause, [], [f382, f721, f1150])).
fof(f382, plain, (~ (e0 = op(e0, e3)) | ~ sP12), inference(cnf_transformation, [], [f133])).
fof(f1155, plain, (~ spl45_118 | ~ spl45_34), inference(avatar_split_clause, [], [f383, f657, f1150])).
fof(f383, plain, (~ (e1 = op(e1, e3)) | ~ sP12), inference(cnf_transformation, [], [f133])).
fof(f1154, plain, (~ spl45_118 | ~ spl45_19), inference(avatar_split_clause, [], [f384, f593, f1150])).
fof(f384, plain, (~ (e2 = op(e2, e3)) | ~ sP12), inference(cnf_transformation, [], [f133])).
fof(f1153, plain, (~ spl45_118 | ~ spl45_4), inference(avatar_split_clause, [], [f385, f529, f1150])).
fof(f385, plain, (~ (e3 = op(e3, e3)) | ~ sP12), inference(cnf_transformation, [], [f133])).
fof(f1148, plain, (~ spl45_117 | spl45_12), inference(avatar_split_clause, [], [f374, f563, f1140])).
fof(f374, plain, ((e3 = op(e3, e1)) | ~ sP13), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP13), inference(nnf_transformation, [], [f68])).
fof(f1147, plain, (~ spl45_117 | spl45_36), inference(avatar_split_clause, [], [f375, f665, f1140])).
fof(f375, plain, ((e3 = op(e1, e3)) | ~ sP13), inference(cnf_transformation, [], [f132])).
fof(f1146, plain, (~ spl45_117 | ~ spl45_49), inference(avatar_split_clause, [], [f376, f721, f1140])).
fof(f376, plain, (~ (e0 = op(e0, e3)) | ~ sP13), inference(cnf_transformation, [], [f132])).
fof(f1145, plain, (~ spl45_117 | ~ spl45_34), inference(avatar_split_clause, [], [f377, f657, f1140])).
fof(f377, plain, (~ (e1 = op(e1, e3)) | ~ sP13), inference(cnf_transformation, [], [f132])).
fof(f1144, plain, (~ spl45_117 | ~ spl45_19), inference(avatar_split_clause, [], [f378, f593, f1140])).
fof(f378, plain, (~ (e2 = op(e2, e3)) | ~ sP13), inference(cnf_transformation, [], [f132])).
fof(f1143, plain, (~ spl45_117 | ~ spl45_4), inference(avatar_split_clause, [], [f379, f529, f1140])).
fof(f379, plain, (~ (e3 = op(e3, e3)) | ~ sP13), inference(cnf_transformation, [], [f132])).
fof(f1138, plain, (~ spl45_116 | spl45_8), inference(avatar_split_clause, [], [f368, f546, f1130])).
fof(f368, plain, ((e3 = op(e3, e2)) | ~ sP14), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e0, e3)) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP14), inference(nnf_transformation, [], [f69])).
fof(f1137, plain, (~ spl45_116 | spl45_20), inference(avatar_split_clause, [], [f369, f597, f1130])).
fof(f369, plain, ((e3 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f131])).
fof(f1136, plain, (~ spl45_116 | ~ spl45_49), inference(avatar_split_clause, [], [f370, f721, f1130])).
fof(f370, plain, (~ (e0 = op(e0, e3)) | ~ sP14), inference(cnf_transformation, [], [f131])).
fof(f1135, plain, (~ spl45_116 | ~ spl45_34), inference(avatar_split_clause, [], [f371, f657, f1130])).
fof(f371, plain, (~ (e1 = op(e1, e3)) | ~ sP14), inference(cnf_transformation, [], [f131])).
fof(f1134, plain, (~ spl45_116 | ~ spl45_19), inference(avatar_split_clause, [], [f372, f593, f1130])).
fof(f372, plain, (~ (e2 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f131])).
fof(f1133, plain, (~ spl45_116 | ~ spl45_4), inference(avatar_split_clause, [], [f373, f529, f1130])).
fof(f373, plain, (~ (e3 = op(e3, e3)) | ~ sP14), inference(cnf_transformation, [], [f131])).
fof(f1128, plain, ~ spl45_114, inference(avatar_split_clause, [], [f501, f1120])).
fof(f501, plain, ~ sP15, inference(trivial_inequality_removal, [], [f366])).
fof(f366, plain, (~ (op(e0, e0) = op(e0, e0)) | ~ sP15), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (((e0 = op(e0, op(e0, e0))) & ~ (op(e0, e0) = op(e0, e0)) & (op(e0, e0) = op(e0, e0))) | ~ sP15), inference(nnf_transformation, [], [f70])).
fof(f1118, plain, (~ spl45_112 | spl45_107), inference(avatar_split_clause, [], [f362, f1085, f1110])).
fof(f362, plain, ((op(e0, e1) = op(e1, e0)) | ~ sP16), inference(cnf_transformation, [], [f129])).
fof(f129, plain, (((e1 = op(e1, op(e0, e1))) & ~ (op(e0, e1) = op(e1, e1)) & (op(e0, e1) = op(e1, e0))) | ~ sP16), inference(nnf_transformation, [], [f71])).
fof(f1117, plain, (~ spl45_112 | spl45_113), inference(avatar_split_clause, [], [f364, f1114, f1110])).
fof(f364, plain, ((e1 = op(e1, op(e0, e1))) | ~ sP16), inference(cnf_transformation, [], [f129])).
fof(f1108, plain, (~ spl45_110 | spl45_98), inference(avatar_split_clause, [], [f359, f1041, f1100])).
fof(f359, plain, ((op(e0, e2) = op(e2, e0)) | ~ sP17), inference(cnf_transformation, [], [f128])).
fof(f128, plain, (((e2 = op(e2, op(e0, e2))) & ~ (op(e0, e2) = op(e2, e2)) & (op(e0, e2) = op(e2, e0))) | ~ sP17), inference(nnf_transformation, [], [f72])).
fof(f1107, plain, (~ spl45_110 | spl45_111), inference(avatar_split_clause, [], [f361, f1104, f1100])).
fof(f361, plain, ((e2 = op(e2, op(e0, e2))) | ~ sP17), inference(cnf_transformation, [], [f128])).
fof(f1098, plain, (~ spl45_108 | spl45_88), inference(avatar_split_clause, [], [f356, f993, f1090])).
fof(f356, plain, ((op(e0, e3) = op(e3, e0)) | ~ sP18), inference(cnf_transformation, [], [f127])).
fof(f127, plain, (((e3 = op(e3, op(e0, e3))) & ~ (op(e0, e3) = op(e3, e3)) & (op(e0, e3) = op(e3, e0))) | ~ sP18), inference(nnf_transformation, [], [f73])).
fof(f1097, plain, (~ spl45_108 | spl45_109), inference(avatar_split_clause, [], [f358, f1094, f1090])).
fof(f358, plain, ((e3 = op(e3, op(e0, e3))) | ~ sP18), inference(cnf_transformation, [], [f127])).
fof(f1088, plain, (~ spl45_105 | spl45_107), inference(avatar_split_clause, [], [f353, f1085, f1076])).
fof(f353, plain, ((op(e0, e1) = op(e1, e0)) | ~ sP19), inference(cnf_transformation, [], [f126])).
fof(f126, plain, (((e0 = op(e0, op(e1, e0))) & ~ (op(e0, e0) = op(e1, e0)) & (op(e0, e1) = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f74])).
fof(f1083, plain, (~ spl45_105 | spl45_106), inference(avatar_split_clause, [], [f355, f1080, f1076])).
fof(f355, plain, ((e0 = op(e0, op(e1, e0))) | ~ sP19), inference(cnf_transformation, [], [f126])).
fof(f1074, plain, ~ spl45_103, inference(avatar_split_clause, [], [f502, f1066])).
fof(f502, plain, ~ sP20, inference(trivial_inequality_removal, [], [f351])).
fof(f351, plain, (~ (op(e1, e1) = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f125])).
fof(f125, plain, (((e1 = op(e1, op(e1, e1))) & ~ (op(e1, e1) = op(e1, e1)) & (op(e1, e1) = op(e1, e1))) | ~ sP20), inference(nnf_transformation, [], [f75])).
fof(f1064, plain, (~ spl45_101 | spl45_95), inference(avatar_split_clause, [], [f347, f1027, f1056])).
fof(f347, plain, ((op(e1, e2) = op(e2, e1)) | ~ sP21), inference(cnf_transformation, [], [f124])).
fof(f124, plain, (((e2 = op(e2, op(e1, e2))) & ~ (op(e1, e2) = op(e2, e2)) & (op(e1, e2) = op(e2, e1))) | ~ sP21), inference(nnf_transformation, [], [f76])).
fof(f1063, plain, (~ spl45_101 | spl45_102), inference(avatar_split_clause, [], [f349, f1060, f1056])).
fof(f349, plain, ((e2 = op(e2, op(e1, e2))) | ~ sP21), inference(cnf_transformation, [], [f124])).
fof(f1054, plain, (~ spl45_99 | spl45_85), inference(avatar_split_clause, [], [f344, f979, f1046])).
fof(f344, plain, ((op(e1, e3) = op(e3, e1)) | ~ sP22), inference(cnf_transformation, [], [f123])).
fof(f123, plain, (((e3 = op(e3, op(e1, e3))) & ~ (op(e1, e3) = op(e3, e3)) & (op(e1, e3) = op(e3, e1))) | ~ sP22), inference(nnf_transformation, [], [f77])).
fof(f1053, plain, (~ spl45_99 | spl45_100), inference(avatar_split_clause, [], [f346, f1050, f1046])).
fof(f346, plain, ((e3 = op(e3, op(e1, e3))) | ~ sP22), inference(cnf_transformation, [], [f123])).
fof(f1044, plain, (~ spl45_96 | spl45_98), inference(avatar_split_clause, [], [f341, f1041, f1032])).
fof(f341, plain, ((op(e0, e2) = op(e2, e0)) | ~ sP23), inference(cnf_transformation, [], [f122])).
fof(f122, plain, (((e0 = op(e0, op(e2, e0))) & ~ (op(e0, e0) = op(e2, e0)) & (op(e0, e2) = op(e2, e0))) | ~ sP23), inference(nnf_transformation, [], [f78])).
fof(f1039, plain, (~ spl45_96 | spl45_97), inference(avatar_split_clause, [], [f343, f1036, f1032])).
fof(f343, plain, ((e0 = op(e0, op(e2, e0))) | ~ sP23), inference(cnf_transformation, [], [f122])).
fof(f1030, plain, (~ spl45_93 | spl45_95), inference(avatar_split_clause, [], [f338, f1027, f1018])).
fof(f338, plain, ((op(e1, e2) = op(e2, e1)) | ~ sP24), inference(cnf_transformation, [], [f121])).
fof(f121, plain, (((e1 = op(e1, op(e2, e1))) & ~ (op(e1, e1) = op(e2, e1)) & (op(e1, e2) = op(e2, e1))) | ~ sP24), inference(nnf_transformation, [], [f79])).
fof(f1025, plain, (~ spl45_93 | spl45_94), inference(avatar_split_clause, [], [f340, f1022, f1018])).
fof(f340, plain, ((e1 = op(e1, op(e2, e1))) | ~ sP24), inference(cnf_transformation, [], [f121])).
fof(f1016, plain, ~ spl45_91, inference(avatar_split_clause, [], [f503, f1008])).
fof(f503, plain, ~ sP25, inference(trivial_inequality_removal, [], [f336])).
fof(f336, plain, (~ (op(e2, e2) = op(e2, e2)) | ~ sP25), inference(cnf_transformation, [], [f120])).
fof(f120, plain, (((e2 = op(e2, op(e2, e2))) & ~ (op(e2, e2) = op(e2, e2)) & (op(e2, e2) = op(e2, e2))) | ~ sP25), inference(nnf_transformation, [], [f80])).
fof(f1006, plain, (~ spl45_89 | spl45_82), inference(avatar_split_clause, [], [f332, f965, f998])).
fof(f332, plain, ((op(e2, e3) = op(e3, e2)) | ~ sP26), inference(cnf_transformation, [], [f119])).
fof(f119, plain, (((e3 = op(e3, op(e2, e3))) & ~ (op(e2, e3) = op(e3, e3)) & (op(e2, e3) = op(e3, e2))) | ~ sP26), inference(nnf_transformation, [], [f81])).
fof(f1005, plain, (~ spl45_89 | spl45_90), inference(avatar_split_clause, [], [f334, f1002, f998])).
fof(f334, plain, ((e3 = op(e3, op(e2, e3))) | ~ sP26), inference(cnf_transformation, [], [f119])).
fof(f996, plain, (~ spl45_86 | spl45_88), inference(avatar_split_clause, [], [f329, f993, f984])).
fof(f329, plain, ((op(e0, e3) = op(e3, e0)) | ~ sP27), inference(cnf_transformation, [], [f118])).
fof(f118, plain, (((e0 = op(e0, op(e3, e0))) & ~ (op(e0, e0) = op(e3, e0)) & (op(e0, e3) = op(e3, e0))) | ~ sP27), inference(nnf_transformation, [], [f82])).
fof(f991, plain, (~ spl45_86 | spl45_87), inference(avatar_split_clause, [], [f331, f988, f984])).
fof(f331, plain, ((e0 = op(e0, op(e3, e0))) | ~ sP27), inference(cnf_transformation, [], [f118])).
fof(f982, plain, (~ spl45_83 | spl45_85), inference(avatar_split_clause, [], [f326, f979, f970])).
fof(f326, plain, ((op(e1, e3) = op(e3, e1)) | ~ sP28), inference(cnf_transformation, [], [f117])).
fof(f117, plain, (((e1 = op(e1, op(e3, e1))) & ~ (op(e1, e1) = op(e3, e1)) & (op(e1, e3) = op(e3, e1))) | ~ sP28), inference(nnf_transformation, [], [f83])).
fof(f977, plain, (~ spl45_83 | spl45_84), inference(avatar_split_clause, [], [f328, f974, f970])).
fof(f328, plain, ((e1 = op(e1, op(e3, e1))) | ~ sP28), inference(cnf_transformation, [], [f117])).
fof(f968, plain, (~ spl45_80 | spl45_82), inference(avatar_split_clause, [], [f323, f965, f956])).
fof(f323, plain, ((op(e2, e3) = op(e3, e2)) | ~ sP29), inference(cnf_transformation, [], [f116])).
fof(f116, plain, (((e2 = op(e2, op(e3, e2))) & ~ (op(e2, e2) = op(e3, e2)) & (op(e2, e3) = op(e3, e2))) | ~ sP29), inference(nnf_transformation, [], [f84])).
fof(f963, plain, (~ spl45_80 | spl45_81), inference(avatar_split_clause, [], [f325, f960, f956])).
fof(f325, plain, ((e2 = op(e2, op(e3, e2))) | ~ sP29), inference(cnf_transformation, [], [f116])).
fof(f954, plain, (~ spl45_79 | spl45_61), inference(avatar_split_clause, [], [f318, f772, f947])).
fof(f318, plain, ((e0 = op(e0, e0)) | ~ sP30), inference(cnf_transformation, [], [f115])).
fof(f115, plain, (((((~ (op(e0, e0) = e3) | ~ (op(e0, e0) = e3)) & (~ (op(e0, e0) = e2) | ~ (op(e0, e0) = e2)) & (~ (op(e0, e0) = e1) | ~ (op(e0, e0) = e1)) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)))) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0))) | ~ sP30), inference(nnf_transformation, [], [f85])).
fof(f953, plain, (~ spl45_79 | ~ spl45_61), inference(avatar_split_clause, [], [f504, f772, f947])).
fof(f504, plain, (~ (e0 = op(e0, e0)) | ~ sP30), inference(duplicate_literal_removal, [], [f319])).
fof(f319, plain, (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ sP30), inference(cnf_transformation, [], [f115])).
fof(f945, plain, (~ spl45_78 | spl45_62), inference(avatar_split_clause, [], [f313, f776, f938])).
fof(f313, plain, ((op(e0, e0) = e1) | ~ sP31), inference(cnf_transformation, [], [f114])).
fof(f114, plain, (((((~ (e3 = op(e1, e0)) | ~ (e3 = op(e0, e1))) & (~ (e2 = op(e1, e0)) | ~ (e2 = op(e0, e1))) & (~ (e1 = op(e1, e0)) | ~ (e1 = op(e0, e1))) & (~ (e0 = op(e1, e0)) | ~ (e0 = op(e0, e1)))) | ~ (op(e0, e0) = e1)) & (op(e0, e0) = e1)) | ~ sP31), inference(nnf_transformation, [], [f86])).
fof(f944, plain, (~ spl45_78 | ~ spl45_62 | ~ spl45_57 | ~ spl45_45), inference(avatar_split_clause, [], [f314, f704, f755, f776, f938])).
fof(f314, plain, (~ (e0 = op(e1, e0)) | ~ (e0 = op(e0, e1)) | ~ (op(e0, e0) = e1) | ~ sP31), inference(cnf_transformation, [], [f114])).
fof(f942, plain, (~ spl45_78 | ~ spl45_62 | ~ spl45_59 | ~ spl45_47), inference(avatar_split_clause, [], [f316, f712, f763, f776, f938])).
fof(f316, plain, (~ (e2 = op(e1, e0)) | ~ (e2 = op(e0, e1)) | ~ (op(e0, e0) = e1) | ~ sP31), inference(cnf_transformation, [], [f114])).
fof(f941, plain, (~ spl45_78 | ~ spl45_62 | ~ spl45_60 | ~ spl45_48), inference(avatar_split_clause, [], [f317, f716, f767, f776, f938])).
fof(f317, plain, (~ (e3 = op(e1, e0)) | ~ (e3 = op(e0, e1)) | ~ (op(e0, e0) = e1) | ~ sP31), inference(cnf_transformation, [], [f114])).
fof(f936, plain, (~ spl45_77 | spl45_63), inference(avatar_split_clause, [], [f308, f780, f929])).
fof(f308, plain, ((op(e0, e0) = e2) | ~ sP32), inference(cnf_transformation, [], [f113])).
fof(f113, plain, (((((~ (e3 = op(e2, e0)) | ~ (e3 = op(e0, e2))) & (~ (e2 = op(e2, e0)) | ~ (e2 = op(e0, e2))) & (~ (e1 = op(e2, e0)) | ~ (e1 = op(e0, e2))) & (~ (e0 = op(e2, e0)) | ~ (e0 = op(e0, e2)))) | ~ (op(e0, e0) = e2)) & (op(e0, e0) = e2)) | ~ sP32), inference(nnf_transformation, [], [f87])).
fof(f935, plain, (~ spl45_77 | ~ spl45_63 | ~ spl45_53 | ~ spl45_29), inference(avatar_split_clause, [], [f309, f636, f738, f780, f929])).
fof(f309, plain, (~ (e0 = op(e2, e0)) | ~ (e0 = op(e0, e2)) | ~ (op(e0, e0) = e2) | ~ sP32), inference(cnf_transformation, [], [f113])).
fof(f934, plain, (~ spl45_77 | ~ spl45_63 | ~ spl45_54 | ~ spl45_30), inference(avatar_split_clause, [], [f310, f640, f742, f780, f929])).
fof(f310, plain, (~ (e1 = op(e2, e0)) | ~ (e1 = op(e0, e2)) | ~ (op(e0, e0) = e2) | ~ sP32), inference(cnf_transformation, [], [f113])).
fof(f932, plain, (~ spl45_77 | ~ spl45_63 | ~ spl45_56 | ~ spl45_32), inference(avatar_split_clause, [], [f312, f648, f750, f780, f929])).
fof(f312, plain, (~ (e3 = op(e2, e0)) | ~ (e3 = op(e0, e2)) | ~ (op(e0, e0) = e2) | ~ sP32), inference(cnf_transformation, [], [f113])).
fof(f927, plain, (~ spl45_76 | spl45_64), inference(avatar_split_clause, [], [f303, f784, f920])).
fof(f303, plain, ((op(e0, e0) = e3) | ~ sP33), inference(cnf_transformation, [], [f112])).
fof(f112, plain, (((((~ (e3 = op(e3, e0)) | ~ (e3 = op(e0, e3))) & (~ (e2 = op(e3, e0)) | ~ (e2 = op(e0, e3))) & (~ (e1 = op(e3, e0)) | ~ (e1 = op(e0, e3))) & (~ (e0 = op(e3, e0)) | ~ (e0 = op(e0, e3)))) | ~ (op(e0, e0) = e3)) & (op(e0, e0) = e3)) | ~ sP33), inference(nnf_transformation, [], [f88])).
fof(f926, plain, (~ spl45_76 | ~ spl45_64 | ~ spl45_49 | ~ spl45_13), inference(avatar_split_clause, [], [f304, f568, f721, f784, f920])).
fof(f304, plain, (~ (e0 = op(e3, e0)) | ~ (e0 = op(e0, e3)) | ~ (op(e0, e0) = e3) | ~ sP33), inference(cnf_transformation, [], [f112])).
fof(f925, plain, (~ spl45_76 | ~ spl45_64 | ~ spl45_50 | ~ spl45_14), inference(avatar_split_clause, [], [f305, f572, f725, f784, f920])).
fof(f305, plain, (~ (e1 = op(e3, e0)) | ~ (e1 = op(e0, e3)) | ~ (op(e0, e0) = e3) | ~ sP33), inference(cnf_transformation, [], [f112])).
fof(f924, plain, (~ spl45_76 | ~ spl45_64 | ~ spl45_51 | ~ spl45_15), inference(avatar_split_clause, [], [f306, f576, f729, f784, f920])).
fof(f306, plain, (~ (e2 = op(e3, e0)) | ~ (e2 = op(e0, e3)) | ~ (op(e0, e0) = e3) | ~ sP33), inference(cnf_transformation, [], [f112])).
fof(f918, plain, (~ spl45_75 | spl45_41), inference(avatar_split_clause, [], [f298, f687, f911])).
fof(f298, plain, ((e0 = op(e1, e1)) | ~ sP34), inference(cnf_transformation, [], [f111])).
fof(f111, plain, (((((~ (e3 = op(e0, e1)) | ~ (e3 = op(e1, e0))) & (~ (e2 = op(e0, e1)) | ~ (e2 = op(e1, e0))) & (~ (e1 = op(e0, e1)) | ~ (e1 = op(e1, e0))) & (~ (e0 = op(e0, e1)) | ~ (e0 = op(e1, e0)))) | ~ (e0 = op(e1, e1))) & (e0 = op(e1, e1))) | ~ sP34), inference(nnf_transformation, [], [f89])).
fof(f916, plain, (~ spl45_75 | ~ spl45_41 | ~ spl45_46 | ~ spl45_58), inference(avatar_split_clause, [], [f300, f759, f708, f687, f911])).
fof(f300, plain, (~ (e1 = op(e0, e1)) | ~ (e1 = op(e1, e0)) | ~ (e0 = op(e1, e1)) | ~ sP34), inference(cnf_transformation, [], [f111])).
fof(f915, plain, (~ spl45_75 | ~ spl45_41 | ~ spl45_47 | ~ spl45_59), inference(avatar_split_clause, [], [f301, f763, f712, f687, f911])).
fof(f301, plain, (~ (e2 = op(e0, e1)) | ~ (e2 = op(e1, e0)) | ~ (e0 = op(e1, e1)) | ~ sP34), inference(cnf_transformation, [], [f111])).
fof(f914, plain, (~ spl45_75 | ~ spl45_41 | ~ spl45_48 | ~ spl45_60), inference(avatar_split_clause, [], [f302, f767, f716, f687, f911])).
fof(f302, plain, (~ (e3 = op(e0, e1)) | ~ (e3 = op(e1, e0)) | ~ (e0 = op(e1, e1)) | ~ sP34), inference(cnf_transformation, [], [f111])).
fof(f909, plain, (~ spl45_74 | spl45_42), inference(avatar_split_clause, [], [f293, f691, f902])).
fof(f293, plain, ((e1 = op(e1, e1)) | ~ sP35), inference(cnf_transformation, [], [f110])).
fof(f110, plain, (((((~ (e3 = op(e1, e1)) | ~ (e3 = op(e1, e1))) & (~ (e2 = op(e1, e1)) | ~ (e2 = op(e1, e1))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e0 = op(e1, e1)))) | ~ (e1 = op(e1, e1))) & (e1 = op(e1, e1))) | ~ sP35), inference(nnf_transformation, [], [f90])).
fof(f907, plain, (~ spl45_74 | ~ spl45_42), inference(avatar_split_clause, [], [f509, f691, f902])).
fof(f509, plain, (~ (e1 = op(e1, e1)) | ~ sP35), inference(duplicate_literal_removal, [], [f295])).
fof(f295, plain, (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ sP35), inference(cnf_transformation, [], [f110])).
fof(f900, plain, (~ spl45_73 | spl45_43), inference(avatar_split_clause, [], [f288, f695, f893])).
fof(f288, plain, ((e2 = op(e1, e1)) | ~ sP36), inference(cnf_transformation, [], [f109])).
fof(f109, plain, (((((~ (e3 = op(e2, e1)) | ~ (e3 = op(e1, e2))) & (~ (e2 = op(e2, e1)) | ~ (e2 = op(e1, e2))) & (~ (e1 = op(e2, e1)) | ~ (e1 = op(e1, e2))) & (~ (e0 = op(e2, e1)) | ~ (e0 = op(e1, e2)))) | ~ (e2 = op(e1, e1))) & (e2 = op(e1, e1))) | ~ sP36), inference(nnf_transformation, [], [f91])).
fof(f899, plain, (~ spl45_73 | ~ spl45_43 | ~ spl45_37 | ~ spl45_25), inference(avatar_split_clause, [], [f289, f619, f670, f695, f893])).
fof(f289, plain, (~ (e0 = op(e2, e1)) | ~ (e0 = op(e1, e2)) | ~ (e2 = op(e1, e1)) | ~ sP36), inference(cnf_transformation, [], [f109])).
fof(f898, plain, (~ spl45_73 | ~ spl45_43 | ~ spl45_38 | ~ spl45_26), inference(avatar_split_clause, [], [f290, f623, f674, f695, f893])).
fof(f290, plain, (~ (e1 = op(e2, e1)) | ~ (e1 = op(e1, e2)) | ~ (e2 = op(e1, e1)) | ~ sP36), inference(cnf_transformation, [], [f109])).
fof(f896, plain, (~ spl45_73 | ~ spl45_43 | ~ spl45_40 | ~ spl45_28), inference(avatar_split_clause, [], [f292, f631, f682, f695, f893])).
fof(f292, plain, (~ (e3 = op(e2, e1)) | ~ (e3 = op(e1, e2)) | ~ (e2 = op(e1, e1)) | ~ sP36), inference(cnf_transformation, [], [f109])).
fof(f891, plain, (~ spl45_72 | spl45_44), inference(avatar_split_clause, [], [f283, f699, f884])).
fof(f283, plain, ((e3 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f108])).
fof(f108, plain, (((((~ (e3 = op(e3, e1)) | ~ (e3 = op(e1, e3))) & (~ (e2 = op(e3, e1)) | ~ (e2 = op(e1, e3))) & (~ (e1 = op(e3, e1)) | ~ (e1 = op(e1, e3))) & (~ (e0 = op(e3, e1)) | ~ (e0 = op(e1, e3)))) | ~ (e3 = op(e1, e1))) & (e3 = op(e1, e1))) | ~ sP37), inference(nnf_transformation, [], [f92])).
fof(f890, plain, (~ spl45_72 | ~ spl45_44 | ~ spl45_33 | ~ spl45_9), inference(avatar_split_clause, [], [f284, f551, f653, f699, f884])).
fof(f284, plain, (~ (e0 = op(e3, e1)) | ~ (e0 = op(e1, e3)) | ~ (e3 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f108])).
fof(f889, plain, (~ spl45_72 | ~ spl45_44 | ~ spl45_34 | ~ spl45_10), inference(avatar_split_clause, [], [f285, f555, f657, f699, f884])).
fof(f285, plain, (~ (e1 = op(e3, e1)) | ~ (e1 = op(e1, e3)) | ~ (e3 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f108])).
fof(f888, plain, (~ spl45_72 | ~ spl45_44 | ~ spl45_35 | ~ spl45_11), inference(avatar_split_clause, [], [f286, f559, f661, f699, f884])).
fof(f286, plain, (~ (e2 = op(e3, e1)) | ~ (e2 = op(e1, e3)) | ~ (e3 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f108])).
fof(f882, plain, (~ spl45_71 | spl45_21), inference(avatar_split_clause, [], [f278, f602, f875])).
fof(f278, plain, ((e0 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f107])).
fof(f107, plain, (((((~ (e3 = op(e0, e2)) | ~ (e3 = op(e2, e0))) & (~ (e2 = op(e0, e2)) | ~ (e2 = op(e2, e0))) & (~ (e1 = op(e0, e2)) | ~ (e1 = op(e2, e0))) & (~ (e0 = op(e0, e2)) | ~ (e0 = op(e2, e0)))) | ~ (e0 = op(e2, e2))) & (e0 = op(e2, e2))) | ~ sP38), inference(nnf_transformation, [], [f93])).
fof(f880, plain, (~ spl45_71 | ~ spl45_21 | ~ spl45_30 | ~ spl45_54), inference(avatar_split_clause, [], [f280, f742, f640, f602, f875])).
fof(f280, plain, (~ (e1 = op(e0, e2)) | ~ (e1 = op(e2, e0)) | ~ (e0 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f107])).
fof(f879, plain, (~ spl45_71 | ~ spl45_21 | ~ spl45_31 | ~ spl45_55), inference(avatar_split_clause, [], [f281, f746, f644, f602, f875])).
fof(f281, plain, (~ (e2 = op(e0, e2)) | ~ (e2 = op(e2, e0)) | ~ (e0 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f107])).
fof(f878, plain, (~ spl45_71 | ~ spl45_21 | ~ spl45_32 | ~ spl45_56), inference(avatar_split_clause, [], [f282, f750, f648, f602, f875])).
fof(f282, plain, (~ (e3 = op(e0, e2)) | ~ (e3 = op(e2, e0)) | ~ (e0 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f107])).
fof(f873, plain, (~ spl45_70 | spl45_22), inference(avatar_split_clause, [], [f273, f606, f866])).
fof(f273, plain, ((e1 = op(e2, e2)) | ~ sP39), inference(cnf_transformation, [], [f106])).
fof(f106, plain, (((((~ (e3 = op(e1, e2)) | ~ (e3 = op(e2, e1))) & (~ (e2 = op(e1, e2)) | ~ (e2 = op(e2, e1))) & (~ (e1 = op(e1, e2)) | ~ (e1 = op(e2, e1))) & (~ (e0 = op(e1, e2)) | ~ (e0 = op(e2, e1)))) | ~ (e1 = op(e2, e2))) & (e1 = op(e2, e2))) | ~ sP39), inference(nnf_transformation, [], [f94])).
fof(f872, plain, (~ spl45_70 | ~ spl45_22 | ~ spl45_25 | ~ spl45_37), inference(avatar_split_clause, [], [f274, f670, f619, f606, f866])).
fof(f274, plain, (~ (e0 = op(e1, e2)) | ~ (e0 = op(e2, e1)) | ~ (e1 = op(e2, e2)) | ~ sP39), inference(cnf_transformation, [], [f106])).
fof(f870, plain, (~ spl45_70 | ~ spl45_22 | ~ spl45_27 | ~ spl45_39), inference(avatar_split_clause, [], [f276, f678, f627, f606, f866])).
fof(f276, plain, (~ (e2 = op(e1, e2)) | ~ (e2 = op(e2, e1)) | ~ (e1 = op(e2, e2)) | ~ sP39), inference(cnf_transformation, [], [f106])).
fof(f869, plain, (~ spl45_70 | ~ spl45_22 | ~ spl45_28 | ~ spl45_40), inference(avatar_split_clause, [], [f277, f682, f631, f606, f866])).
fof(f277, plain, (~ (e3 = op(e1, e2)) | ~ (e3 = op(e2, e1)) | ~ (e1 = op(e2, e2)) | ~ sP39), inference(cnf_transformation, [], [f106])).
fof(f864, plain, (~ spl45_69 | spl45_23), inference(avatar_split_clause, [], [f268, f610, f857])).
fof(f268, plain, ((e2 = op(e2, e2)) | ~ sP40), inference(cnf_transformation, [], [f105])).
fof(f105, plain, (((((~ (e3 = op(e2, e2)) | ~ (e3 = op(e2, e2))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e1 = op(e2, e2))) & (~ (e0 = op(e2, e2)) | ~ (e0 = op(e2, e2)))) | ~ (e2 = op(e2, e2))) & (e2 = op(e2, e2))) | ~ sP40), inference(nnf_transformation, [], [f95])).
fof(f861, plain, (~ spl45_69 | ~ spl45_23), inference(avatar_split_clause, [], [f514, f610, f857])).
fof(f514, plain, (~ (e2 = op(e2, e2)) | ~ sP40), inference(duplicate_literal_removal, [], [f271])).
fof(f271, plain, (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ sP40), inference(cnf_transformation, [], [f105])).
fof(f855, plain, (~ spl45_68 | spl45_24), inference(avatar_split_clause, [], [f263, f614, f848])).
fof(f263, plain, ((e3 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f104])).
fof(f104, plain, (((((~ (e3 = op(e3, e2)) | ~ (e3 = op(e2, e3))) & (~ (e2 = op(e3, e2)) | ~ (e2 = op(e2, e3))) & (~ (e1 = op(e3, e2)) | ~ (e1 = op(e2, e3))) & (~ (e0 = op(e3, e2)) | ~ (e0 = op(e2, e3)))) | ~ (e3 = op(e2, e2))) & (e3 = op(e2, e2))) | ~ sP41), inference(nnf_transformation, [], [f96])).
fof(f854, plain, (~ spl45_68 | ~ spl45_24 | ~ spl45_17 | ~ spl45_5), inference(avatar_split_clause, [], [f264, f534, f585, f614, f848])).
fof(f264, plain, (~ (e0 = op(e3, e2)) | ~ (e0 = op(e2, e3)) | ~ (e3 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f104])).
fof(f853, plain, (~ spl45_68 | ~ spl45_24 | ~ spl45_18 | ~ spl45_6), inference(avatar_split_clause, [], [f265, f538, f589, f614, f848])).
fof(f265, plain, (~ (e1 = op(e3, e2)) | ~ (e1 = op(e2, e3)) | ~ (e3 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f104])).
fof(f852, plain, (~ spl45_68 | ~ spl45_24 | ~ spl45_19 | ~ spl45_7), inference(avatar_split_clause, [], [f266, f542, f593, f614, f848])).
fof(f266, plain, (~ (e2 = op(e3, e2)) | ~ (e2 = op(e2, e3)) | ~ (e3 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f104])).
fof(f846, plain, (~ spl45_67 | spl45_1), inference(avatar_split_clause, [], [f258, f517, f839])).
fof(f258, plain, ((e0 = op(e3, e3)) | ~ sP42), inference(cnf_transformation, [], [f103])).
fof(f103, plain, (((((~ (e3 = op(e0, e3)) | ~ (e3 = op(e3, e0))) & (~ (e2 = op(e0, e3)) | ~ (e2 = op(e3, e0))) & (~ (e1 = op(e0, e3)) | ~ (e1 = op(e3, e0))) & (~ (e0 = op(e0, e3)) | ~ (e0 = op(e3, e0)))) | ~ (e0 = op(e3, e3))) & (e0 = op(e3, e3))) | ~ sP42), inference(nnf_transformation, [], [f97])).
fof(f843, plain, (~ spl45_67 | ~ spl45_1 | ~ spl45_15 | ~ spl45_51), inference(avatar_split_clause, [], [f261, f729, f576, f517, f839])).
fof(f261, plain, (~ (e2 = op(e0, e3)) | ~ (e2 = op(e3, e0)) | ~ (e0 = op(e3, e3)) | ~ sP42), inference(cnf_transformation, [], [f103])).
fof(f842, plain, (~ spl45_67 | ~ spl45_1 | ~ spl45_16 | ~ spl45_52), inference(avatar_split_clause, [], [f262, f733, f580, f517, f839])).
fof(f262, plain, (~ (e3 = op(e0, e3)) | ~ (e3 = op(e3, e0)) | ~ (e0 = op(e3, e3)) | ~ sP42), inference(cnf_transformation, [], [f103])).
fof(f837, plain, (~ spl45_66 | spl45_2), inference(avatar_split_clause, [], [f253, f521, f830])).
fof(f253, plain, ((e1 = op(e3, e3)) | ~ sP43), inference(cnf_transformation, [], [f102])).
fof(f102, plain, (((((~ (e3 = op(e1, e3)) | ~ (e3 = op(e3, e1))) & (~ (e2 = op(e1, e3)) | ~ (e2 = op(e3, e1))) & (~ (e1 = op(e1, e3)) | ~ (e1 = op(e3, e1))) & (~ (e0 = op(e1, e3)) | ~ (e0 = op(e3, e1)))) | ~ (e1 = op(e3, e3))) & (e1 = op(e3, e3))) | ~ sP43), inference(nnf_transformation, [], [f98])).
fof(f834, plain, (~ spl45_66 | ~ spl45_2 | ~ spl45_11 | ~ spl45_35), inference(avatar_split_clause, [], [f256, f661, f559, f521, f830])).
fof(f256, plain, (~ (e2 = op(e1, e3)) | ~ (e2 = op(e3, e1)) | ~ (e1 = op(e3, e3)) | ~ sP43), inference(cnf_transformation, [], [f102])).
fof(f833, plain, (~ spl45_66 | ~ spl45_2 | ~ spl45_12 | ~ spl45_36), inference(avatar_split_clause, [], [f257, f665, f563, f521, f830])).
fof(f257, plain, (~ (e3 = op(e1, e3)) | ~ (e3 = op(e3, e1)) | ~ (e1 = op(e3, e3)) | ~ sP43), inference(cnf_transformation, [], [f102])).
fof(f828, plain, (~ spl45_65 | spl45_3), inference(avatar_split_clause, [], [f248, f525, f821])).
fof(f248, plain, ((e2 = op(e3, e3)) | ~ sP44), inference(cnf_transformation, [], [f101])).
fof(f101, plain, (((((~ (e3 = op(e2, e3)) | ~ (e3 = op(e3, e2))) & (~ (e2 = op(e2, e3)) | ~ (e2 = op(e3, e2))) & (~ (e1 = op(e2, e3)) | ~ (e1 = op(e3, e2))) & (~ (e0 = op(e2, e3)) | ~ (e0 = op(e3, e2)))) | ~ (e2 = op(e3, e3))) & (e2 = op(e3, e3))) | ~ sP44), inference(nnf_transformation, [], [f99])).
fof(f827, plain, (~ spl45_65 | ~ spl45_3 | ~ spl45_5 | ~ spl45_17), inference(avatar_split_clause, [], [f249, f585, f534, f525, f821])).
fof(f249, plain, (~ (e0 = op(e2, e3)) | ~ (e0 = op(e3, e2)) | ~ (e2 = op(e3, e3)) | ~ sP44), inference(cnf_transformation, [], [f101])).
fof(f826, plain, (~ spl45_65 | ~ spl45_3 | ~ spl45_6 | ~ spl45_18), inference(avatar_split_clause, [], [f250, f589, f538, f525, f821])).
fof(f250, plain, (~ (e1 = op(e2, e3)) | ~ (e1 = op(e3, e2)) | ~ (e2 = op(e3, e3)) | ~ sP44), inference(cnf_transformation, [], [f101])).
fof(f824, plain, (~ spl45_65 | ~ spl45_3 | ~ spl45_8 | ~ spl45_20), inference(avatar_split_clause, [], [f252, f597, f546, f525, f821])).
fof(f252, plain, (~ (e3 = op(e2, e3)) | ~ (e3 = op(e3, e2)) | ~ (e2 = op(e3, e3)) | ~ sP44), inference(cnf_transformation, [], [f101])).
fof(f819, plain, (spl45_61 | spl45_57 | spl45_53 | spl45_49), inference(avatar_split_clause, [], [f162, f721, f738, f755, f772])).
fof(f162, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax2)).
fof(f818, plain, (spl45_61 | spl45_45 | spl45_29 | spl45_13), inference(avatar_split_clause, [], [f163, f568, f636, f704, f772])).
fof(f163, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f817, plain, (spl45_62 | spl45_58 | spl45_54 | spl45_50), inference(avatar_split_clause, [], [f164, f725, f742, f759, f776])).
fof(f164, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f816, plain, (spl45_62 | spl45_46 | spl45_30 | spl45_14), inference(avatar_split_clause, [], [f165, f572, f640, f708, f776])).
fof(f165, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f815, plain, (spl45_63 | spl45_59 | spl45_55 | spl45_51), inference(avatar_split_clause, [], [f166, f729, f746, f763, f780])).
fof(f166, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f814, plain, (spl45_63 | spl45_47 | spl45_31 | spl45_15), inference(avatar_split_clause, [], [f167, f576, f644, f712, f780])).
fof(f167, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f813, plain, (spl45_64 | spl45_60 | spl45_56 | spl45_52), inference(avatar_split_clause, [], [f168, f733, f750, f767, f784])).
fof(f168, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f812, plain, (spl45_64 | spl45_48 | spl45_32 | spl45_16), inference(avatar_split_clause, [], [f169, f580, f648, f716, f784])).
fof(f169, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f811, plain, (spl45_45 | spl45_41 | spl45_37 | spl45_33), inference(avatar_split_clause, [], [f170, f653, f670, f687, f704])).
fof(f170, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f810, plain, (spl45_57 | spl45_41 | spl45_25 | spl45_9), inference(avatar_split_clause, [], [f171, f551, f619, f687, f755])).
fof(f171, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f809, plain, (spl45_46 | spl45_42 | spl45_38 | spl45_34), inference(avatar_split_clause, [], [f172, f657, f674, f691, f708])).
fof(f172, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f808, plain, (spl45_58 | spl45_42 | spl45_26 | spl45_10), inference(avatar_split_clause, [], [f173, f555, f623, f691, f759])).
fof(f173, plain, ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f807, plain, (spl45_47 | spl45_43 | spl45_39 | spl45_35), inference(avatar_split_clause, [], [f174, f661, f678, f695, f712])).
fof(f174, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f806, plain, (spl45_59 | spl45_43 | spl45_27 | spl45_11), inference(avatar_split_clause, [], [f175, f559, f627, f695, f763])).
fof(f175, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f805, plain, (spl45_48 | spl45_44 | spl45_40 | spl45_36), inference(avatar_split_clause, [], [f176, f665, f682, f699, f716])).
fof(f176, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f804, plain, (spl45_60 | spl45_44 | spl45_28 | spl45_12), inference(avatar_split_clause, [], [f177, f563, f631, f699, f767])).
fof(f177, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f803, plain, (spl45_29 | spl45_25 | spl45_21 | spl45_17), inference(avatar_split_clause, [], [f178, f585, f602, f619, f636])).
fof(f178, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f802, plain, (spl45_53 | spl45_37 | spl45_21 | spl45_5), inference(avatar_split_clause, [], [f179, f534, f602, f670, f738])).
fof(f179, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f801, plain, (spl45_30 | spl45_26 | spl45_22 | spl45_18), inference(avatar_split_clause, [], [f180, f589, f606, f623, f640])).
fof(f180, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f800, plain, (spl45_54 | spl45_38 | spl45_22 | spl45_6), inference(avatar_split_clause, [], [f181, f538, f606, f674, f742])).
fof(f181, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f799, plain, (spl45_31 | spl45_27 | spl45_23 | spl45_19), inference(avatar_split_clause, [], [f182, f593, f610, f627, f644])).
fof(f182, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f798, plain, (spl45_55 | spl45_39 | spl45_23 | spl45_7), inference(avatar_split_clause, [], [f183, f542, f610, f678, f746])).
fof(f183, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f797, plain, (spl45_32 | spl45_28 | spl45_24 | spl45_20), inference(avatar_split_clause, [], [f184, f597, f614, f631, f648])).
fof(f184, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f796, plain, (spl45_56 | spl45_40 | spl45_24 | spl45_8), inference(avatar_split_clause, [], [f185, f546, f614, f682, f750])).
fof(f185, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f795, plain, (spl45_13 | spl45_9 | spl45_5 | spl45_1), inference(avatar_split_clause, [], [f186, f517, f534, f551, f568])).
fof(f186, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f794, plain, (spl45_49 | spl45_33 | spl45_17 | spl45_1), inference(avatar_split_clause, [], [f187, f517, f585, f653, f721])).
fof(f187, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f793, plain, (spl45_14 | spl45_10 | spl45_6 | spl45_2), inference(avatar_split_clause, [], [f188, f521, f538, f555, f572])).
fof(f188, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f792, plain, (spl45_50 | spl45_34 | spl45_18 | spl45_2), inference(avatar_split_clause, [], [f189, f521, f589, f657, f725])).
fof(f189, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f791, plain, (spl45_15 | spl45_11 | spl45_7 | spl45_3), inference(avatar_split_clause, [], [f190, f525, f542, f559, f576])).
fof(f190, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f790, plain, (spl45_51 | spl45_35 | spl45_19 | spl45_3), inference(avatar_split_clause, [], [f191, f525, f593, f661, f729])).
fof(f191, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f789, plain, (spl45_16 | spl45_12 | spl45_8 | spl45_4), inference(avatar_split_clause, [], [f192, f529, f546, f563, f580])).
fof(f192, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f788, plain, (spl45_52 | spl45_36 | spl45_20 | spl45_4), inference(avatar_split_clause, [], [f193, f529, f597, f665, f733])).
fof(f193, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f787, plain, (spl45_61 | spl45_62 | spl45_63 | spl45_64), inference(avatar_split_clause, [], [f146, f784, f780, f776, f772])).
fof(f146, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG134+1.p', ax1)).
fof(f770, plain, (spl45_57 | spl45_58 | spl45_59 | spl45_60), inference(avatar_split_clause, [], [f147, f767, f763, f759, f755])).
fof(f147, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f753, plain, (spl45_53 | spl45_54 | spl45_55 | spl45_56), inference(avatar_split_clause, [], [f148, f750, f746, f742, f738])).
fof(f148, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f736, plain, (spl45_49 | spl45_50 | spl45_51 | spl45_52), inference(avatar_split_clause, [], [f149, f733, f729, f725, f721])).
fof(f149, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f719, plain, (spl45_45 | spl45_46 | spl45_47 | spl45_48), inference(avatar_split_clause, [], [f150, f716, f712, f708, f704])).
fof(f150, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f702, plain, (spl45_41 | spl45_42 | spl45_43 | spl45_44), inference(avatar_split_clause, [], [f151, f699, f695, f691, f687])).
fof(f151, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f685, plain, (spl45_37 | spl45_38 | spl45_39 | spl45_40), inference(avatar_split_clause, [], [f152, f682, f678, f674, f670])).
fof(f152, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f668, plain, (spl45_33 | spl45_34 | spl45_35 | spl45_36), inference(avatar_split_clause, [], [f153, f665, f661, f657, f653])).
fof(f153, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f651, plain, (spl45_29 | spl45_30 | spl45_31 | spl45_32), inference(avatar_split_clause, [], [f154, f648, f644, f640, f636])).
fof(f154, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f634, plain, (spl45_25 | spl45_26 | spl45_27 | spl45_28), inference(avatar_split_clause, [], [f155, f631, f627, f623, f619])).
fof(f155, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f617, plain, (spl45_21 | spl45_22 | spl45_23 | spl45_24), inference(avatar_split_clause, [], [f156, f614, f610, f606, f602])).
fof(f156, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f600, plain, (spl45_17 | spl45_18 | spl45_19 | spl45_20), inference(avatar_split_clause, [], [f157, f597, f593, f589, f585])).
fof(f157, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f583, plain, (spl45_13 | spl45_14 | spl45_15 | spl45_16), inference(avatar_split_clause, [], [f158, f580, f576, f572, f568])).
fof(f158, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f566, plain, (spl45_9 | spl45_10 | spl45_11 | spl45_12), inference(avatar_split_clause, [], [f159, f563, f559, f555, f551])).
fof(f159, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f549, plain, (spl45_5 | spl45_6 | spl45_7 | spl45_8), inference(avatar_split_clause, [], [f160, f546, f542, f538, f534])).
fof(f160, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f532, plain, (spl45_1 | spl45_2 | spl45_3 | spl45_4), inference(avatar_split_clause, [], [f161, f529, f525, f521, f517])).
fof(f161, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).