fof(f5644, plain, $false, inference(avatar_sat_refutation, [], [f1227, f1244, f1295, f1312, f1329, f1397, f1414, f1431, f1448, f1466, f1467, f1477, f1478, f1479, f1480, f1484, f1485, f1486, f1487, f1489, f1493, f1494, f1495, f1531, f1548, f1582, f1599, f1616, f1633, f1667, f1701, f1718, f1735, f1752, f1771, f1772, f1777, f1779, f1780, f1782, f1783, f1784, f1785, f1786, f1787, f1788, f1791, f1798, f1799, f1800, f1801, f1828, f1833, f1861, f1882, f1890, f1898, f1906, f1912, f1913, f1914, f1919, f1920, f1922, f1927, f1928, f1930, f1936, f1945, f1953, f1961, f1969, f1978, f1986, f1994, f2002, f2008, f2016, f2017, f2024, f2026, f2031, f2033, f2034, f2041, f2042, f2049, f2050, f2056, f2057, f2058, f2064, f2071, f2073, f2074, f2079, f2080, f2082, f2087, f2088, f2096, f2106, f2114, f2122, f2130, f2138, f2146, f2154, f2162, f2167, f2168, f2169, f2170, f2175, f2176, f2178, f2184, f2185, f2192, f2201, f2202, f2209, f2210, f2217, f2218, f2224, f2232, f2240, f2242, f2248, f2249, f2255, f2256, f2258, f2265, f2273, f2281, f2289, f2295, f2296, f2298, f2303, f2304, f2306, f2311, f2312, f2314, f2320, f2327, f2330, f2335, f2336, f2338, f2343, f2344, f2352, f2360, f2369, f2370, f2377, f2378, f2385, f2386, f2409, f2413, f2419, f2425, f2435, f2440, f2445, f2473, f2479, f2489, f2494, f2502, f2510, f2518, f2524, f2531, f2532, f2534, f2539, f2540, f2542, f2548, f2557, f2565, f2573, f2581, f2590, f2598, f2606, f2614, f2620, f2628, f2629, f2630, f2636, f2637, f2638, f2643, f2645, f2646, f2652, f2661, f2662, f2669, f2670, f2676, f2683, f2685, f2686, f2692, f2693, f2699, f2700, f2708, f2716, f2724, f2725, f2726, f2731, f2732, f2734, f2740, f2750, f2758, f2766, f2774, f2780, f2787, f2788, f2796, f2797, f2804, f2813, f2814, f2821, f2822, f2829, f2830, f2836, f2844, f2851, f2852, f2854, f2860, f2861, f2868, f2877, f2885, f2893, f2901, f2908, f2915, f2916, f2918, f2923, f2924, f2926, f2932, f2939, f2942, f2947, f2948, f2950, f2955, f2956, f2964, f2974, f2982, f2990, f2998, f3000, f3011, f3021, f3025, f3026, f3027, f3256, f3270, f3284, f3720, f3726, f3735, f3736, f3745, f3746, f3747, f3749, f3770, f3773, f3776, f3781, f3812, f3816, f3818, f3838, f3861, f3873, f3874, f3875, f3876, f3877, f3878, f3879, f3880, f3889, f3891, f3900, f3901, f3908, f3909, f3922, f3934, f3949, f3957, f3965, f3978, f3984, f4014, f4017, f4019, f4025, f4061, f4071, f4087, f4088, f4089, f4090, f4093, f4102, f4116, f4134, f4146, f4156, f4160, f4168, f4172, f4188, f4199, f4206, f4211, f4212, f4213, f4226, f4228, f4232, f4241, f4245, f4255, f4274, f4296, f4302, f4317, f4336, f4343, f4363, f4400, f4409, f4450, f4461, f4477, f4495, f4500, f4502, f4510, f4515, f4520, f4535, f4541, f4562, f4575, f4585, f4642, f4644, f4646, f4653, f4701, f4709, f4719, f4730, f4734, f4746, f4758, f4764, f4780, f4796, f4821, f4833, f4853, f4862, f4863, f4868, f4884, f4913, f4925, f4930, f4935, f4945, f4950, f4959, f4964, f4969, f4978, f4982, f4990, f4994, f5011, f5021, f5031, f5038, f5082, f5096, f5097, f5111, f5124, f5127, f5136, f5163, f5181, f5195, f5198, f5207, f5213, f5218, f5228, f5236, f5237, f5245, f5252, f5253, f5302, f5305, f5320, f5325, f5335, f5376, f5381, f5384, f5412, f5415, f5470, f5484, f5492, f5518, f5530, f5552, f5585, f5605, f5615, f5622, f5623, f5630, f5635])).
fof(f5635, plain, (~ spl144_19 | ~ spl144_23), inference(avatar_split_clause, [], [f5631, f1288, f1271])).
fof(f1271, plain, (spl144_19 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_19])])).
fof(f1288, plain, (spl144_23 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_23])])).
fof(f5631, plain, (~ (e12 = op1(e12, e13)) | ~ spl144_23), inference(backward_demodulation, [], [f449, f1290])).
fof(f1290, plain, ((e12 = op1(e12, e12)) | ~ spl144_23), inference(avatar_component_clause, [], [f1288])).
fof(f449, plain, ~ (op1(e12, e12) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax5)).
fof(f5630, plain, (~ spl144_26 | ~ spl144_27), inference(avatar_contradiction_clause, [], [f5629])).
fof(f5629, plain, ($false | (~ spl144_26 | ~ spl144_27)), inference(subsumption_resolution, [], [f5628, f507])).
fof(f507, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax7)).
fof(f5628, plain, ((e11 = e12) | (~ spl144_26 | ~ spl144_27)), inference(backward_demodulation, [], [f1307, f1303])).
fof(f1303, plain, ((e11 = op1(e12, e11)) | ~ spl144_26), inference(avatar_component_clause, [], [f1301])).
fof(f1301, plain, (spl144_26 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_26])])).
fof(f1307, plain, ((e12 = op1(e12, e11)) | ~ spl144_27), inference(avatar_component_clause, [], [f1305])).
fof(f1305, plain, (spl144_27 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_27])])).
fof(f5623, plain, (~ spl144_19 | ~ spl144_27), inference(avatar_split_clause, [], [f5618, f1305, f1271])).
fof(f5618, plain, (~ (e12 = op1(e12, e13)) | ~ spl144_27), inference(backward_demodulation, [], [f448, f1307])).
fof(f448, plain, ~ (op1(e12, e11) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f5622, plain, (~ spl144_23 | ~ spl144_27), inference(avatar_split_clause, [], [f5617, f1305, f1288])).
fof(f5617, plain, (~ (e12 = op1(e12, e12)) | ~ spl144_27), inference(backward_demodulation, [], [f447, f1307])).
fof(f447, plain, ~ (op1(e12, e11) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f5615, plain, (~ spl144_19 | ~ spl144_35), inference(avatar_split_clause, [], [f5612, f1339, f1271])).
fof(f1339, plain, (spl144_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_35])])).
fof(f5612, plain, (~ (e12 = op1(e12, e13)) | ~ spl144_35), inference(backward_demodulation, [], [f429, f1341])).
fof(f1341, plain, ((e12 = op1(e11, e13)) | ~ spl144_35), inference(avatar_component_clause, [], [f1339])).
fof(f429, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f5605, plain, (~ spl144_36 | ~ spl144_44), inference(avatar_split_clause, [], [f5601, f1377, f1343])).
fof(f1343, plain, (spl144_36 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_36])])).
fof(f1377, plain, (spl144_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_44])])).
fof(f5601, plain, (~ (e13 = op1(e11, e13)) | ~ spl144_44), inference(backward_demodulation, [], [f442, f1379])).
fof(f1379, plain, ((e13 = op1(e11, e11)) | ~ spl144_44), inference(avatar_component_clause, [], [f1377])).
fof(f442, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f5585, plain, (~ spl144_43 | ~ spl144_59), inference(avatar_split_clause, [], [f5578, f1441, f1373])).
fof(f1373, plain, (spl144_43 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_43])])).
fof(f1441, plain, (spl144_59 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_59])])).
fof(f5578, plain, (~ (e12 = op1(e11, e11)) | ~ spl144_59), inference(backward_demodulation, [], [f414, f1443])).
fof(f1443, plain, ((e12 = op1(e10, e11)) | ~ spl144_59), inference(avatar_component_clause, [], [f1441])).
fof(f414, plain, ~ (op1(e10, e11) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f5552, plain, (~ spl144_99 | ~ spl144_318), inference(avatar_split_clause, [], [f5551, f3190, f1643])).
fof(f1643, plain, (spl144_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_99])])).
fof(f3190, plain, (spl144_318 <=> (e22 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_318])])).
fof(f5551, plain, (~ (e22 = op2(e21, e23)) | ~ spl144_318), inference(forward_demodulation, [], [f3045, f3191])).
fof(f3191, plain, ((e22 = h2(e13)) | ~ spl144_318), inference(avatar_component_clause, [], [f3190])).
fof(f3045, plain, ~ (op2(e21, e23) = h2(e13)), inference(backward_demodulation, [], [f490, f1121])).
fof(f1121, plain, (op2(e21, e21) = h2(e13)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((op2(e21, e21) = h2(e13)) & (h2(e12) = op2(op2(e21, e21), e21)) & (h2(e11) = op2(op2(e21, e21), op2(e21, e21))) & (e21 = h2(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax15)).
fof(f490, plain, ~ (op2(e21, e21) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax6)).
fof(f5530, plain, (~ spl144_66 | ~ spl144_68), inference(avatar_contradiction_clause, [], [f5529])).
fof(f5529, plain, ($false | (~ spl144_66 | ~ spl144_68)), inference(subsumption_resolution, [], [f5528, f514])).
fof(f514, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax8)).
fof(f5528, plain, ((e21 = e23) | (~ spl144_66 | ~ spl144_68)), inference(forward_demodulation, [], [f1513, f1505])).
fof(f1505, plain, ((e21 = op2(e23, e23)) | ~ spl144_66), inference(avatar_component_clause, [], [f1503])).
fof(f1503, plain, (spl144_66 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_66])])).
fof(f1513, plain, ((e23 = op2(e23, e23)) | ~ spl144_68), inference(avatar_component_clause, [], [f1511])).
fof(f1511, plain, (spl144_68 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_68])])).
fof(f5518, plain, (~ spl144_76 | ~ spl144_92), inference(avatar_split_clause, [], [f5517, f1613, f1545])).
fof(f1545, plain, (spl144_76 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_76])])).
fof(f1613, plain, (spl144_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_92])])).
fof(f5517, plain, (~ (e23 = op2(e23, e21)) | ~ spl144_92), inference(backward_demodulation, [], [f467, f1615])).
fof(f1615, plain, ((e23 = op2(e22, e21)) | ~ spl144_92), inference(avatar_component_clause, [], [f1613])).
fof(f467, plain, ~ (op2(e22, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f5492, plain, (spl144_113 | ~ spl144_128 | ~ spl144_225), inference(avatar_split_clause, [], [f5491, f2491, f1766, f1703])).
fof(f1703, plain, (spl144_113 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_113])])).
fof(f1766, plain, (spl144_128 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl144_128])])).
fof(f2491, plain, (spl144_225 <=> (e20 = op2(e20, op2(e20, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl144_225])])).
fof(f5491, plain, ((e20 = op2(e20, e23)) | (~ spl144_128 | ~ spl144_225)), inference(forward_demodulation, [], [f2493, f1768])).
fof(f1768, plain, ((op2(e20, e20) = e23) | ~ spl144_128), inference(avatar_component_clause, [], [f1766])).
fof(f2493, plain, ((e20 = op2(e20, op2(e20, e20))) | ~ spl144_225), inference(avatar_component_clause, [], [f2491])).
fof(f5484, plain, (~ spl144_90 | ~ spl144_310), inference(avatar_split_clause, [], [f5477, f3150, f1605])).
fof(f1605, plain, (spl144_90 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_90])])).
fof(f3150, plain, (spl144_310 <=> (e21 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_310])])).
fof(f5477, plain, (~ (e21 = op2(e22, e21)) | ~ spl144_310), inference(backward_demodulation, [], [f3052, f3151])).
fof(f3151, plain, ((e21 = h3(e13)) | ~ spl144_310), inference(avatar_component_clause, [], [f3150])).
fof(f3052, plain, ~ (op2(e22, e21) = h3(e13)), inference(backward_demodulation, [], [f495, f1125])).
fof(f1125, plain, (op2(e22, e22) = h3(e13)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(e22, e22) = h3(e13)) & (h3(e12) = op2(op2(e22, e22), e22)) & (h3(e11) = op2(op2(e22, e22), op2(e22, e22))) & (e22 = h3(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax16)).
fof(f495, plain, ~ (op2(e22, e21) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f5470, plain, (~ spl144_123 | ~ spl144_318), inference(avatar_split_clause, [], [f5461, f3190, f1745])).
fof(f1745, plain, (spl144_123 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_123])])).
fof(f5461, plain, (~ (e22 = op2(e20, e21)) | ~ spl144_318), inference(backward_demodulation, [], [f3040, f3191])).
fof(f3040, plain, ~ (op2(e20, e21) = h2(e13)), inference(backward_demodulation, [], [f462, f1121])).
fof(f462, plain, ~ (op2(e20, e21) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f5415, plain, (~ spl144_72 | ~ spl144_104), inference(avatar_split_clause, [], [f5414, f1664, f1528])).
fof(f1528, plain, (spl144_72 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_72])])).
fof(f1664, plain, (spl144_104 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_104])])).
fof(f5414, plain, (~ (e23 = op2(e23, e22)) | ~ spl144_104), inference(forward_demodulation, [], [f472, f1666])).
fof(f1666, plain, ((e23 = op2(e21, e22)) | ~ spl144_104), inference(avatar_component_clause, [], [f1664])).
fof(f472, plain, ~ (op2(e21, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f5412, plain, (~ spl144_83 | ~ spl144_84), inference(avatar_contradiction_clause, [], [f5411])).
fof(f5411, plain, ($false | (~ spl144_83 | ~ spl144_84)), inference(subsumption_resolution, [], [f5410, f515])).
fof(f515, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f5410, plain, ((e22 = e23) | (~ spl144_83 | ~ spl144_84)), inference(forward_demodulation, [], [f1581, f1577])).
fof(f1577, plain, ((e22 = op2(e22, e23)) | ~ spl144_83), inference(avatar_component_clause, [], [f1575])).
fof(f1575, plain, (spl144_83 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_83])])).
fof(f1581, plain, ((e23 = op2(e22, e23)) | ~ spl144_84), inference(avatar_component_clause, [], [f1579])).
fof(f1579, plain, (spl144_84 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_84])])).
fof(f5384, plain, (~ spl144_118 | ~ spl144_119), inference(avatar_contradiction_clause, [], [f5383])).
fof(f5383, plain, ($false | (~ spl144_118 | ~ spl144_119)), inference(subsumption_resolution, [], [f5382, f513])).
fof(f513, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f8])).
fof(f5382, plain, ((e21 = e22) | (~ spl144_118 | ~ spl144_119)), inference(backward_demodulation, [], [f1730, f1726])).
fof(f1726, plain, ((e21 = op2(e20, e22)) | ~ spl144_118), inference(avatar_component_clause, [], [f1724])).
fof(f1724, plain, (spl144_118 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_118])])).
fof(f1730, plain, ((e22 = op2(e20, e22)) | ~ spl144_119), inference(avatar_component_clause, [], [f1728])).
fof(f1728, plain, (spl144_119 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_119])])).
fof(f5381, plain, (~ spl144_125 | ~ spl144_128), inference(avatar_contradiction_clause, [], [f5380])).
fof(f5380, plain, ($false | (~ spl144_125 | ~ spl144_128)), inference(subsumption_resolution, [], [f5379, f512])).
fof(f512, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f8])).
fof(f5379, plain, ((e20 = e23) | (~ spl144_125 | ~ spl144_128)), inference(backward_demodulation, [], [f1768, f1756])).
fof(f1756, plain, ((e20 = op2(e20, e20)) | ~ spl144_125), inference(avatar_component_clause, [], [f1754])).
fof(f1754, plain, (spl144_125 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_125])])).
fof(f5376, plain, (~ spl144_88 | ~ spl144_83 | spl144_212), inference(avatar_split_clause, [], [f5375, f2422, f1575, f1596])).
fof(f1596, plain, (spl144_88 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_88])])).
fof(f2422, plain, (spl144_212 <=> (e23 = op2(e22, op2(e22, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl144_212])])).
fof(f5375, plain, (~ (e23 = op2(e22, e22)) | (~ spl144_83 | spl144_212)), inference(forward_demodulation, [], [f2423, f1577])).
fof(f2423, plain, (~ (e23 = op2(e22, op2(e22, e23))) | spl144_212), inference(avatar_component_clause, [], [f2422])).
fof(f5335, plain, (~ spl144_314 | ~ spl144_69), inference(avatar_split_clause, [], [f5334, f1516, f3170])).
fof(f3170, plain, (spl144_314 <=> (e20 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_314])])).
fof(f1516, plain, (spl144_69 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_69])])).
fof(f5334, plain, (~ (e20 = h3(e13)) | ~ spl144_69), inference(forward_demodulation, [], [f3050, f1518])).
fof(f1518, plain, ((e20 = op2(e23, e22)) | ~ spl144_69), inference(avatar_component_clause, [], [f1516])).
fof(f3050, plain, ~ (op2(e23, e22) = h3(e13)), inference(backward_demodulation, [], [f473, f1125])).
fof(f473, plain, ~ (op2(e22, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f5325, plain, (~ spl144_314 | ~ spl144_93), inference(avatar_split_clause, [], [f5143, f1618, f3170])).
fof(f1618, plain, (spl144_93 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_93])])).
fof(f5143, plain, (~ (e20 = h3(e13)) | ~ spl144_93), inference(forward_demodulation, [], [f3051, f1620])).
fof(f1620, plain, ((e20 = op2(e22, e20)) | ~ spl144_93), inference(avatar_component_clause, [], [f1618])).
fof(f3051, plain, ~ (op2(e22, e20) = h3(e13)), inference(backward_demodulation, [], [f493, f1125])).
fof(f493, plain, ~ (op2(e22, e20) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f5320, plain, (spl144_118 | ~ spl144_123 | ~ spl144_224), inference(avatar_split_clause, [], [f5319, f2486, f1745, f1724])).
fof(f2486, plain, (spl144_224 <=> (e21 = op2(e20, op2(e20, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl144_224])])).
fof(f5319, plain, ((e21 = op2(e20, e22)) | (~ spl144_123 | ~ spl144_224)), inference(forward_demodulation, [], [f2488, f1747])).
fof(f1747, plain, ((e22 = op2(e20, e21)) | ~ spl144_123), inference(avatar_component_clause, [], [f1745])).
fof(f2488, plain, ((e21 = op2(e20, op2(e20, e21))) | ~ spl144_224), inference(avatar_component_clause, [], [f2486])).
fof(f5305, plain, (~ spl144_114 | ~ spl144_298), inference(avatar_split_clause, [], [f4497, f3089, f1707])).
fof(f1707, plain, (spl144_114 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_114])])).
fof(f3089, plain, (spl144_298 <=> (e21 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_298])])).
fof(f4497, plain, (~ (e21 = op2(e20, e23)) | ~ spl144_298), inference(forward_demodulation, [], [f3056, f3090])).
fof(f3090, plain, ((e21 = h4(e13)) | ~ spl144_298), inference(avatar_component_clause, [], [f3089])).
fof(f3056, plain, ~ (op2(e20, e23) = h4(e13)), inference(backward_demodulation, [], [f476, f1129])).
fof(f1129, plain, (op2(e23, e23) = h4(e13)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(e23, e23) = h4(e13)) & (h4(e12) = op2(op2(e23, e23), e23)) & (h4(e11) = op2(op2(e23, e23), op2(e23, e23))) & (e23 = h4(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax17)).
fof(f476, plain, ~ (op2(e20, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f5302, plain, (~ spl144_115 | ~ spl144_119), inference(avatar_split_clause, [], [f5301, f1728, f1711])).
fof(f1711, plain, (spl144_115 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_115])])).
fof(f5301, plain, (~ (e22 = op2(e20, e23)) | ~ spl144_119), inference(forward_demodulation, [], [f485, f1730])).
fof(f485, plain, ~ (op2(e20, e22) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f5253, plain, (~ spl144_42 | ~ spl144_46), inference(avatar_split_clause, [], [f5251, f1386, f1369])).
fof(f1369, plain, (spl144_42 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_42])])).
fof(f1386, plain, (spl144_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_46])])).
fof(f5251, plain, (~ (e11 = op1(e11, e11)) | ~ spl144_46), inference(backward_demodulation, [], [f438, f1388])).
fof(f1388, plain, ((e11 = op1(e11, e10)) | ~ spl144_46), inference(avatar_component_clause, [], [f1386])).
fof(f438, plain, ~ (op1(e11, e10) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f5252, plain, (~ spl144_30 | ~ spl144_46), inference(avatar_split_clause, [], [f5248, f1386, f1318])).
fof(f1318, plain, (spl144_30 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_30])])).
fof(f5248, plain, (~ (e11 = op1(e12, e10)) | ~ spl144_46), inference(backward_demodulation, [], [f411, f1388])).
fof(f411, plain, ~ (op1(e11, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f5245, plain, (~ spl144_39 | ~ spl144_55), inference(avatar_split_clause, [], [f5239, f1424, f1356])).
fof(f1356, plain, (spl144_39 <=> (e12 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_39])])).
fof(f1424, plain, (spl144_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_55])])).
fof(f5239, plain, (~ (e12 = op1(e11, e12)) | ~ spl144_55), inference(backward_demodulation, [], [f420, f1426])).
fof(f1426, plain, ((e12 = op1(e10, e12)) | ~ spl144_55), inference(avatar_component_clause, [], [f1424])).
fof(f420, plain, ~ (op1(e10, e12) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f5237, plain, (~ spl144_54 | ~ spl144_58), inference(avatar_split_clause, [], [f5230, f1437, f1420])).
fof(f1420, plain, (spl144_54 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_54])])).
fof(f1437, plain, (spl144_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_58])])).
fof(f5230, plain, (~ (e11 = op1(e10, e12)) | ~ spl144_58), inference(backward_demodulation, [], [f435, f1439])).
fof(f1439, plain, ((e11 = op1(e10, e11)) | ~ spl144_58), inference(avatar_component_clause, [], [f1437])).
fof(f435, plain, ~ (op1(e10, e11) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f5236, plain, (~ spl144_42 | ~ spl144_58), inference(avatar_split_clause, [], [f5229, f1437, f1369])).
fof(f5229, plain, (~ (e11 = op1(e11, e11)) | ~ spl144_58), inference(backward_demodulation, [], [f414, f1439])).
fof(f5228, plain, (~ spl144_66 | ~ spl144_67), inference(avatar_contradiction_clause, [], [f5227])).
fof(f5227, plain, ($false | (~ spl144_66 | ~ spl144_67)), inference(subsumption_resolution, [], [f5226, f513])).
fof(f5226, plain, ((e21 = e22) | (~ spl144_66 | ~ spl144_67)), inference(forward_demodulation, [], [f1509, f1505])).
fof(f1509, plain, ((e22 = op2(e23, e23)) | ~ spl144_67), inference(avatar_component_clause, [], [f1507])).
fof(f1507, plain, (spl144_67 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_67])])).
fof(f5218, plain, (~ spl144_72 | ~ spl144_76), inference(avatar_split_clause, [], [f5215, f1545, f1528])).
fof(f5215, plain, (~ (e23 = op2(e23, e22)) | ~ spl144_76), inference(backward_demodulation, [], [f501, f1547])).
fof(f1547, plain, ((e23 = op2(e23, e21)) | ~ spl144_76), inference(avatar_component_clause, [], [f1545])).
fof(f501, plain, ~ (op2(e23, e21) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f5213, plain, (~ spl144_90 | spl144_214), inference(avatar_contradiction_clause, [], [f5212])).
fof(f5212, plain, ($false | (~ spl144_90 | spl144_214)), inference(subsumption_resolution, [], [f5210, f1607])).
fof(f1607, plain, ((e21 = op2(e22, e21)) | ~ spl144_90), inference(avatar_component_clause, [], [f1605])).
fof(f5210, plain, (~ (e21 = op2(e22, e21)) | (~ spl144_90 | spl144_214)), inference(backward_demodulation, [], [f2433, f1607])).
fof(f2433, plain, (~ (e21 = op2(e22, op2(e22, e21))) | spl144_214), inference(avatar_component_clause, [], [f2432])).
fof(f2432, plain, (spl144_214 <=> (e21 = op2(e22, op2(e22, e21)))), introduced(avatar_definition, [new_symbols(naming, [spl144_214])])).
fof(f5207, plain, (~ spl144_102 | ~ spl144_103), inference(avatar_contradiction_clause, [], [f5206])).
fof(f5206, plain, ($false | (~ spl144_102 | ~ spl144_103)), inference(subsumption_resolution, [], [f5205, f513])).
fof(f5205, plain, ((e21 = e22) | (~ spl144_102 | ~ spl144_103)), inference(backward_demodulation, [], [f1662, f1658])).
fof(f1658, plain, ((e21 = op2(e21, e22)) | ~ spl144_102), inference(avatar_component_clause, [], [f1656])).
fof(f1656, plain, (spl144_102 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_102])])).
fof(f1662, plain, ((e22 = op2(e21, e22)) | ~ spl144_103), inference(avatar_component_clause, [], [f1660])).
fof(f1660, plain, (spl144_103 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_103])])).
fof(f5198, plain, (~ spl144_122 | ~ spl144_123), inference(avatar_contradiction_clause, [], [f5197])).
fof(f5197, plain, ($false | (~ spl144_122 | ~ spl144_123)), inference(subsumption_resolution, [], [f5196, f513])).
fof(f5196, plain, ((e21 = e22) | (~ spl144_122 | ~ spl144_123)), inference(forward_demodulation, [], [f1747, f1743])).
fof(f1743, plain, ((e21 = op2(e20, e21)) | ~ spl144_122), inference(avatar_component_clause, [], [f1741])).
fof(f1741, plain, (spl144_122 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_122])])).
fof(f5195, plain, (spl144_88 | ~ spl144_83 | ~ spl144_212), inference(avatar_split_clause, [], [f5194, f2422, f1575, f1596])).
fof(f5194, plain, ((e23 = op2(e22, e22)) | (~ spl144_83 | ~ spl144_212)), inference(forward_demodulation, [], [f2424, f1577])).
fof(f2424, plain, ((e23 = op2(e22, op2(e22, e23))) | ~ spl144_212), inference(avatar_component_clause, [], [f2422])).
fof(f5181, plain, (~ spl144_73 | ~ spl144_326), inference(avatar_split_clause, [], [f5173, f3231, f1533])).
fof(f1533, plain, (spl144_73 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_73])])).
fof(f3231, plain, (spl144_326 <=> (e20 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_326])])).
fof(f5173, plain, (~ (e20 = op2(e23, e21)) | ~ spl144_326), inference(backward_demodulation, [], [f3042, f3232])).
fof(f3232, plain, ((e20 = h2(e13)) | ~ spl144_326), inference(avatar_component_clause, [], [f3231])).
fof(f3042, plain, ~ (op2(e23, e21) = h2(e13)), inference(backward_demodulation, [], [f466, f1121])).
fof(f466, plain, ~ (op2(e21, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f5163, plain, (~ spl144_72 | ~ spl144_362), inference(avatar_split_clause, [], [f5155, f3449, f1528])).
fof(f3449, plain, (spl144_362 <=> (e23 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_362])])).
fof(f5155, plain, (~ (e23 = op2(e23, e22)) | ~ spl144_362), inference(backward_demodulation, [], [f3050, f3450])).
fof(f3450, plain, ((e23 = h3(e13)) | ~ spl144_362), inference(avatar_component_clause, [], [f3449])).
fof(f5136, plain, (~ spl144_102 | ~ spl144_110), inference(avatar_split_clause, [], [f4817, f1690, f1656])).
fof(f1690, plain, (spl144_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_110])])).
fof(f4817, plain, (~ (e21 = op2(e21, e22)) | ~ spl144_110), inference(backward_demodulation, [], [f487, f1692])).
fof(f1692, plain, ((e21 = op2(e21, e20)) | ~ spl144_110), inference(avatar_component_clause, [], [f1690])).
fof(f487, plain, ~ (op2(e21, e20) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f5127, plain, (~ spl144_2 | ~ spl144_4), inference(avatar_contradiction_clause, [], [f5126])).
fof(f5126, plain, ($false | (~ spl144_2 | ~ spl144_4)), inference(subsumption_resolution, [], [f5125, f508])).
fof(f508, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f7])).
fof(f5125, plain, ((e11 = e13) | (~ spl144_2 | ~ spl144_4)), inference(forward_demodulation, [], [f1209, f1201])).
fof(f1201, plain, ((e11 = op1(e13, e13)) | ~ spl144_2), inference(avatar_component_clause, [], [f1199])).
fof(f1199, plain, (spl144_2 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_2])])).
fof(f1209, plain, ((e13 = op1(e13, e13)) | ~ spl144_4), inference(avatar_component_clause, [], [f1207])).
fof(f1207, plain, (spl144_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_4])])).
fof(f5124, plain, (~ spl144_5 | ~ spl144_72 | ~ spl144_392 | ~ spl144_394), inference(avatar_contradiction_clause, [], [f5123])).
fof(f5123, plain, ($false | (~ spl144_5 | ~ spl144_72 | ~ spl144_392 | ~ spl144_394)), inference(subsumption_resolution, [], [f5122, f512])).
fof(f5122, plain, ((e20 = e23) | (~ spl144_5 | ~ spl144_72 | ~ spl144_392 | ~ spl144_394)), inference(forward_demodulation, [], [f5119, f1114])).
fof(f1114, plain, (e20 = h1(e10)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((op2(e20, e20) = h1(e13)) & (op2(op2(e20, e20), e20) = h1(e12)) & (op2(op2(e20, e20), op2(e20, e20)) = h1(e11)) & (e20 = h1(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax14)).
fof(f5119, plain, ((e23 = h1(e10)) | (~ spl144_5 | ~ spl144_72 | ~ spl144_392 | ~ spl144_394)), inference(backward_demodulation, [], [f5023, f1214])).
fof(f1214, plain, ((e10 = op1(e13, e12)) | ~ spl144_5), inference(avatar_component_clause, [], [f1212])).
fof(f1212, plain, (spl144_5 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_5])])).
fof(f5023, plain, ((e23 = h1(op1(e13, e12))) | (~ spl144_72 | ~ spl144_392 | ~ spl144_394)), inference(forward_demodulation, [], [f4936, f1530])).
fof(f1530, plain, ((e23 = op2(e23, e22)) | ~ spl144_72), inference(avatar_component_clause, [], [f1528])).
fof(f4936, plain, ((op2(e23, e22) = h1(op1(e13, e12))) | (~ spl144_392 | ~ spl144_394)), inference(forward_demodulation, [], [f3662, f3654])).
fof(f3654, plain, ((e23 = h1(e13)) | ~ spl144_392), inference(avatar_component_clause, [], [f3653])).
fof(f3653, plain, (spl144_392 <=> (e23 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_392])])).
fof(f3662, plain, ((h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ spl144_394), inference(avatar_component_clause, [], [f3661])).
fof(f3661, plain, (spl144_394 <=> (h1(op1(e13, e12)) = op2(h1(e13), e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_394])])).
fof(f5111, plain, (~ spl144_8 | ~ spl144_24), inference(avatar_split_clause, [], [f5108, f1292, f1224])).
fof(f1224, plain, (spl144_8 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_8])])).
fof(f1292, plain, (spl144_24 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_24])])).
fof(f5108, plain, (~ (e13 = op1(e13, e12)) | ~ spl144_24), inference(backward_demodulation, [], [f425, f1294])).
fof(f1294, plain, ((e13 = op1(e12, e12)) | ~ spl144_24), inference(avatar_component_clause, [], [f1292])).
fof(f425, plain, ~ (op1(e12, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f5097, plain, (~ spl144_22 | ~ spl144_30), inference(avatar_split_clause, [], [f5094, f1318, f1284])).
fof(f1284, plain, (spl144_22 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_22])])).
fof(f5094, plain, (~ (e11 = op1(e12, e12)) | ~ spl144_30), inference(backward_demodulation, [], [f445, f1320])).
fof(f1320, plain, ((e11 = op1(e12, e10)) | ~ spl144_30), inference(avatar_component_clause, [], [f1318])).
fof(f445, plain, ~ (op1(e12, e10) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f5096, plain, (spl144_25 | ~ spl144_30 | ~ spl144_133), inference(avatar_split_clause, [], [f5092, f1825, f1318, f1297])).
fof(f1297, plain, (spl144_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_25])])).
fof(f1825, plain, (spl144_133 <=> (e10 = op1(e12, op1(e12, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl144_133])])).
fof(f5092, plain, ((e10 = op1(e12, e11)) | (~ spl144_30 | ~ spl144_133)), inference(backward_demodulation, [], [f1827, f1320])).
fof(f1827, plain, ((e10 = op1(e12, op1(e12, e10))) | ~ spl144_133), inference(avatar_component_clause, [], [f1825])).
fof(f5082, plain, (~ spl144_37 | ~ spl144_45), inference(avatar_split_clause, [], [f5079, f1382, f1348])).
fof(f1348, plain, (spl144_37 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_37])])).
fof(f1382, plain, (spl144_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_45])])).
fof(f5079, plain, (~ (e10 = op1(e11, e12)) | ~ spl144_45), inference(backward_demodulation, [], [f439, f1384])).
fof(f1384, plain, ((e10 = op1(e11, e10)) | ~ spl144_45), inference(avatar_component_clause, [], [f1382])).
fof(f439, plain, ~ (op1(e11, e10) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f5038, plain, (~ spl144_38 | ~ spl144_42), inference(avatar_split_clause, [], [f5037, f1369, f1352])).
fof(f1352, plain, (spl144_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_38])])).
fof(f5037, plain, (~ (e11 = op1(e11, e12)) | ~ spl144_42), inference(forward_demodulation, [], [f441, f1371])).
fof(f1371, plain, ((e11 = op1(e11, e11)) | ~ spl144_42), inference(avatar_component_clause, [], [f1369])).
fof(f441, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f5031, plain, (~ spl144_25 | ~ spl144_9), inference(avatar_split_clause, [], [f4448, f1229, f1297])).
fof(f1229, plain, (spl144_9 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_9])])).
fof(f4448, plain, (~ (e10 = op1(e12, e11)) | ~ spl144_9), inference(forward_demodulation, [], [f419, f1231])).
fof(f1231, plain, ((e10 = op1(e13, e11)) | ~ spl144_9), inference(avatar_component_clause, [], [f1229])).
fof(f419, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f5021, plain, (~ spl144_42 | ~ spl144_43), inference(avatar_contradiction_clause, [], [f5020])).
fof(f5020, plain, ($false | (~ spl144_42 | ~ spl144_43)), inference(subsumption_resolution, [], [f5019, f507])).
fof(f5019, plain, ((e11 = e12) | (~ spl144_42 | ~ spl144_43)), inference(backward_demodulation, [], [f1375, f1371])).
fof(f1375, plain, ((e12 = op1(e11, e11)) | ~ spl144_43), inference(avatar_component_clause, [], [f1373])).
fof(f5011, plain, (~ spl144_62 | ~ spl144_64), inference(avatar_contradiction_clause, [], [f5010])).
fof(f5010, plain, ($false | (~ spl144_62 | ~ spl144_64)), inference(subsumption_resolution, [], [f5009, f508])).
fof(f5009, plain, ((e11 = e13) | (~ spl144_62 | ~ spl144_64)), inference(backward_demodulation, [], [f1464, f1456])).
fof(f1456, plain, ((op1(e10, e10) = e11) | ~ spl144_62), inference(avatar_component_clause, [], [f1454])).
fof(f1454, plain, (spl144_62 <=> (op1(e10, e10) = e11)), introduced(avatar_definition, [new_symbols(naming, [spl144_62])])).
fof(f1464, plain, ((op1(e10, e10) = e13) | ~ spl144_64), inference(avatar_component_clause, [], [f1462])).
fof(f1462, plain, (spl144_64 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl144_64])])).
fof(f4994, plain, (~ spl144_55 | ~ spl144_119 | spl144_406), inference(avatar_contradiction_clause, [], [f4993])).
fof(f4993, plain, ($false | (~ spl144_55 | ~ spl144_119 | spl144_406)), inference(subsumption_resolution, [], [f4992, f1730])).
fof(f4992, plain, (~ (e22 = op2(e20, e22)) | (~ spl144_55 | spl144_406)), inference(forward_demodulation, [], [f4991, f3037])).
fof(f3037, plain, (e22 = h1(e12)), inference(forward_demodulation, [], [f3036, f3035])).
fof(f3035, plain, (e22 = op2(h1(e13), e20)), inference(backward_demodulation, [], [f1112, f1117])).
fof(f1117, plain, (op2(e20, e20) = h1(e13)), inference(cnf_transformation, [], [f14])).
fof(f1112, plain, (e22 = op2(op2(e20, e20), e20)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((op2(e20, e20) = e23) & (e22 = op2(op2(e20, e20), e20)) & (e21 = op2(op2(e20, e20), op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax13)).
fof(f3036, plain, (h1(e12) = op2(h1(e13), e20)), inference(forward_demodulation, [], [f1116, f1117])).
fof(f1116, plain, (op2(op2(e20, e20), e20) = h1(e12)), inference(cnf_transformation, [], [f14])).
fof(f4991, plain, (~ (op2(e20, e22) = h1(e12)) | (~ spl144_55 | spl144_406)), inference(forward_demodulation, [], [f3711, f1426])).
fof(f3711, plain, (~ (op2(e20, e22) = h1(op1(e10, e12))) | spl144_406), inference(avatar_component_clause, [], [f3709])).
fof(f3709, plain, (spl144_406 <=> (op2(e20, e22) = h1(op1(e10, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl144_406])])).
fof(f4990, plain, (~ spl144_29 | ~ spl144_93 | spl144_400), inference(avatar_contradiction_clause, [], [f4989])).
fof(f4989, plain, ($false | (~ spl144_29 | ~ spl144_93 | spl144_400)), inference(subsumption_resolution, [], [f4988, f1620])).
fof(f4988, plain, (~ (e20 = op2(e22, e20)) | (~ spl144_29 | spl144_400)), inference(forward_demodulation, [], [f4987, f1114])).
fof(f4987, plain, (~ (op2(e22, e20) = h1(e10)) | (~ spl144_29 | spl144_400)), inference(forward_demodulation, [], [f3687, f1316])).
fof(f1316, plain, ((e10 = op1(e12, e10)) | ~ spl144_29), inference(avatar_component_clause, [], [f1314])).
fof(f1314, plain, (spl144_29 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_29])])).
fof(f3687, plain, (~ (op2(e22, e20) = h1(op1(e12, e10))) | spl144_400), inference(avatar_component_clause, [], [f3685])).
fof(f3685, plain, (spl144_400 <=> (op2(e22, e20) = h1(op1(e12, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl144_400])])).
fof(f4982, plain, (~ spl144_37 | ~ spl144_101 | spl144_402), inference(avatar_contradiction_clause, [], [f4981])).
fof(f4981, plain, ($false | (~ spl144_37 | ~ spl144_101 | spl144_402)), inference(subsumption_resolution, [], [f4980, f1654])).
fof(f1654, plain, ((e20 = op2(e21, e22)) | ~ spl144_101), inference(avatar_component_clause, [], [f1652])).
fof(f1652, plain, (spl144_101 <=> (e20 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_101])])).
fof(f4980, plain, (~ (e20 = op2(e21, e22)) | (~ spl144_37 | spl144_402)), inference(forward_demodulation, [], [f4979, f1114])).
fof(f4979, plain, (~ (op2(e21, e22) = h1(e10)) | (~ spl144_37 | spl144_402)), inference(forward_demodulation, [], [f3695, f1350])).
fof(f1350, plain, ((e10 = op1(e11, e12)) | ~ spl144_37), inference(avatar_component_clause, [], [f1348])).
fof(f3695, plain, (~ (op2(e21, e22) = h1(op1(e11, e12))) | spl144_402), inference(avatar_component_clause, [], [f3693])).
fof(f3693, plain, (spl144_402 <=> (op2(e21, e22) = h1(op1(e11, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl144_402])])).
fof(f4978, plain, (~ spl144_22 | ~ spl144_310 | spl144_398), inference(avatar_contradiction_clause, [], [f4977])).
fof(f4977, plain, ($false | (~ spl144_22 | ~ spl144_310 | spl144_398)), inference(subsumption_resolution, [], [f4976, f3039])).
fof(f3039, plain, (e21 = h1(e11)), inference(forward_demodulation, [], [f3038, f3034])).
fof(f3034, plain, (e21 = op2(h1(e13), h1(e13))), inference(backward_demodulation, [], [f1111, f1117])).
fof(f1111, plain, (e21 = op2(op2(e20, e20), op2(e20, e20))), inference(cnf_transformation, [], [f13])).
fof(f3038, plain, (h1(e11) = op2(h1(e13), h1(e13))), inference(forward_demodulation, [], [f1115, f1117])).
fof(f1115, plain, (op2(op2(e20, e20), op2(e20, e20)) = h1(e11)), inference(cnf_transformation, [], [f14])).
fof(f4976, plain, (~ (e21 = h1(e11)) | (~ spl144_22 | ~ spl144_310 | spl144_398)), inference(forward_demodulation, [], [f4975, f3151])).
fof(f4975, plain, (~ (h1(e11) = h3(e13)) | (~ spl144_22 | spl144_398)), inference(forward_demodulation, [], [f3679, f1286])).
fof(f1286, plain, ((e11 = op1(e12, e12)) | ~ spl144_22), inference(avatar_component_clause, [], [f1284])).
fof(f3679, plain, (~ (h3(e13) = h1(op1(e12, e12))) | spl144_398), inference(avatar_component_clause, [], [f3677])).
fof(f3677, plain, (spl144_398 <=> (h3(e13) = h1(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl144_398])])).
fof(f4969, plain, (~ spl144_46 | ~ spl144_110 | spl144_404), inference(avatar_contradiction_clause, [], [f4968])).
fof(f4968, plain, ($false | (~ spl144_46 | ~ spl144_110 | spl144_404)), inference(subsumption_resolution, [], [f4967, f1692])).
fof(f4967, plain, (~ (e21 = op2(e21, e20)) | (~ spl144_46 | spl144_404)), inference(forward_demodulation, [], [f4966, f3039])).
fof(f4966, plain, (~ (op2(e21, e20) = h1(e11)) | (~ spl144_46 | spl144_404)), inference(forward_demodulation, [], [f3703, f1388])).
fof(f3703, plain, (~ (op2(e21, e20) = h1(op1(e11, e10))) | spl144_404), inference(avatar_component_clause, [], [f3701])).
fof(f3701, plain, (spl144_404 <=> (op2(e21, e20) = h1(op1(e11, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl144_404])])).
fof(f4964, plain, (~ spl144_19 | ~ spl144_83 | ~ spl144_392 | spl144_397), inference(avatar_contradiction_clause, [], [f4963])).
fof(f4963, plain, ($false | (~ spl144_19 | ~ spl144_83 | ~ spl144_392 | spl144_397)), inference(subsumption_resolution, [], [f4962, f1577])).
fof(f4962, plain, (~ (e22 = op2(e22, e23)) | (~ spl144_19 | ~ spl144_392 | spl144_397)), inference(forward_demodulation, [], [f4961, f3037])).
fof(f4961, plain, (~ (op2(e22, e23) = h1(e12)) | (~ spl144_19 | ~ spl144_392 | spl144_397)), inference(forward_demodulation, [], [f4960, f1273])).
fof(f1273, plain, ((e12 = op1(e12, e13)) | ~ spl144_19), inference(avatar_component_clause, [], [f1271])).
fof(f4960, plain, (~ (op2(e22, e23) = h1(op1(e12, e13))) | (~ spl144_392 | spl144_397)), inference(forward_demodulation, [], [f3675, f3654])).
fof(f3675, plain, (~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | spl144_397), inference(avatar_component_clause, [], [f3673])).
fof(f3673, plain, (spl144_397 <=> (h1(op1(e12, e13)) = op2(e22, h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl144_397])])).
fof(f4959, plain, (~ spl144_15 | spl144_396), inference(avatar_contradiction_clause, [], [f4958])).
fof(f4958, plain, ($false | (~ spl144_15 | spl144_396)), inference(subsumption_resolution, [], [f4957, f3037])).
fof(f4957, plain, (~ (e22 = h1(e12)) | (~ spl144_15 | spl144_396)), inference(forward_demodulation, [], [f3671, f1256])).
fof(f1256, plain, ((e12 = op1(e13, e10)) | ~ spl144_15), inference(avatar_component_clause, [], [f1254])).
fof(f1254, plain, (spl144_15 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_15])])).
fof(f3671, plain, (~ (e22 = h1(op1(e13, e10))) | spl144_396), inference(avatar_component_clause, [], [f3669])).
fof(f3669, plain, (spl144_396 <=> (e22 = h1(op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl144_396])])).
fof(f4950, plain, (~ spl144_49 | ~ spl144_113 | ~ spl144_392 | spl144_405), inference(avatar_contradiction_clause, [], [f4949])).
fof(f4949, plain, ($false | (~ spl144_49 | ~ spl144_113 | ~ spl144_392 | spl144_405)), inference(subsumption_resolution, [], [f4948, f1705])).
fof(f1705, plain, ((e20 = op2(e20, e23)) | ~ spl144_113), inference(avatar_component_clause, [], [f1703])).
fof(f4948, plain, (~ (e20 = op2(e20, e23)) | (~ spl144_49 | ~ spl144_392 | spl144_405)), inference(forward_demodulation, [], [f4947, f1114])).
fof(f4947, plain, (~ (op2(e20, e23) = h1(e10)) | (~ spl144_49 | ~ spl144_392 | spl144_405)), inference(forward_demodulation, [], [f4946, f1401])).
fof(f1401, plain, ((e10 = op1(e10, e13)) | ~ spl144_49), inference(avatar_component_clause, [], [f1399])).
fof(f1399, plain, (spl144_49 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_49])])).
fof(f4946, plain, (~ (op2(e20, e23) = h1(op1(e10, e13))) | (~ spl144_392 | spl144_405)), inference(forward_demodulation, [], [f3707, f3654])).
fof(f3707, plain, (~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | spl144_405), inference(avatar_component_clause, [], [f3705])).
fof(f3705, plain, (spl144_405 <=> (h1(op1(e10, e13)) = op2(e20, h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl144_405])])).
fof(f4945, plain, (~ spl144_9 | ~ spl144_73 | ~ spl144_392 | spl144_395), inference(avatar_contradiction_clause, [], [f4944])).
fof(f4944, plain, ($false | (~ spl144_9 | ~ spl144_73 | ~ spl144_392 | spl144_395)), inference(subsumption_resolution, [], [f4943, f1535])).
fof(f1535, plain, ((e20 = op2(e23, e21)) | ~ spl144_73), inference(avatar_component_clause, [], [f1533])).
fof(f4943, plain, (~ (e20 = op2(e23, e21)) | (~ spl144_9 | ~ spl144_392 | spl144_395)), inference(forward_demodulation, [], [f4942, f1114])).
fof(f4942, plain, (~ (op2(e23, e21) = h1(e10)) | (~ spl144_9 | ~ spl144_392 | spl144_395)), inference(forward_demodulation, [], [f4941, f1231])).
fof(f4941, plain, (~ (op2(e23, e21) = h1(op1(e13, e11))) | (~ spl144_392 | spl144_395)), inference(forward_demodulation, [], [f3667, f3654])).
fof(f3667, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | spl144_395), inference(avatar_component_clause, [], [f3665])).
fof(f3665, plain, (spl144_395 <=> (h1(op1(e13, e11)) = op2(h1(e13), e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_395])])).
fof(f4935, plain, (~ spl144_58 | ~ spl144_122 | spl144_407), inference(avatar_contradiction_clause, [], [f4934])).
fof(f4934, plain, ($false | (~ spl144_58 | ~ spl144_122 | spl144_407)), inference(subsumption_resolution, [], [f4933, f1743])).
fof(f4933, plain, (~ (e21 = op2(e20, e21)) | (~ spl144_58 | spl144_407)), inference(forward_demodulation, [], [f4932, f3039])).
fof(f4932, plain, (~ (op2(e20, e21) = h1(e11)) | (~ spl144_58 | spl144_407)), inference(forward_demodulation, [], [f3715, f1439])).
fof(f3715, plain, (~ (op2(e20, e21) = h1(op1(e10, e11))) | spl144_407), inference(avatar_component_clause, [], [f3713])).
fof(f3713, plain, (spl144_407 <=> (op2(e20, e21) = h1(op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl144_407])])).
fof(f4930, plain, (~ spl144_8 | ~ spl144_72 | ~ spl144_392 | spl144_394), inference(avatar_contradiction_clause, [], [f4929])).
fof(f4929, plain, ($false | (~ spl144_8 | ~ spl144_72 | ~ spl144_392 | spl144_394)), inference(subsumption_resolution, [], [f4928, f1530])).
fof(f4928, plain, (~ (e23 = op2(e23, e22)) | (~ spl144_8 | ~ spl144_392 | spl144_394)), inference(forward_demodulation, [], [f4927, f3654])).
fof(f4927, plain, (~ (op2(e23, e22) = h1(e13)) | (~ spl144_8 | ~ spl144_392 | spl144_394)), inference(forward_demodulation, [], [f4926, f1226])).
fof(f1226, plain, ((e13 = op1(e13, e12)) | ~ spl144_8), inference(avatar_component_clause, [], [f1224])).
fof(f4926, plain, (~ (op2(e23, e22) = h1(op1(e13, e12))) | (~ spl144_392 | spl144_394)), inference(forward_demodulation, [], [f3663, f3654])).
fof(f3663, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | spl144_394), inference(avatar_component_clause, [], [f3661])).
fof(f4925, plain, (~ spl144_2 | spl144_393), inference(avatar_contradiction_clause, [], [f4924])).
fof(f4924, plain, ($false | (~ spl144_2 | spl144_393)), inference(subsumption_resolution, [], [f4923, f3039])).
fof(f4923, plain, (~ (e21 = h1(e11)) | (~ spl144_2 | spl144_393)), inference(forward_demodulation, [], [f3659, f1201])).
fof(f3659, plain, (~ (e21 = h1(op1(e13, e13))) | spl144_393), inference(avatar_component_clause, [], [f3657])).
fof(f3657, plain, (spl144_393 <=> (e21 = h1(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl144_393])])).
fof(f4913, plain, (~ spl144_64 | spl144_408), inference(avatar_contradiction_clause, [], [f4912])).
fof(f4912, plain, ($false | (~ spl144_64 | spl144_408)), inference(trivial_inequality_removal, [], [f4911])).
fof(f4911, plain, (~ (h1(e13) = h1(e13)) | (~ spl144_64 | spl144_408)), inference(forward_demodulation, [], [f3719, f1464])).
fof(f3719, plain, (~ (h1(e13) = h1(op1(e10, e10))) | spl144_408), inference(avatar_component_clause, [], [f3717])).
fof(f3717, plain, (spl144_408 <=> (h1(e13) = h1(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl144_408])])).
fof(f4884, plain, (~ spl144_43 | ~ spl144_318 | spl144_403), inference(avatar_contradiction_clause, [], [f4883])).
fof(f4883, plain, ($false | (~ spl144_43 | ~ spl144_318 | spl144_403)), inference(subsumption_resolution, [], [f4882, f3037])).
fof(f4882, plain, (~ (e22 = h1(e12)) | (~ spl144_43 | ~ spl144_318 | spl144_403)), inference(forward_demodulation, [], [f4881, f3191])).
fof(f4881, plain, (~ (h1(e12) = h2(e13)) | (~ spl144_43 | spl144_403)), inference(forward_demodulation, [], [f3699, f1375])).
fof(f3699, plain, (~ (h2(e13) = h1(op1(e11, e11))) | spl144_403), inference(avatar_component_clause, [], [f3697])).
fof(f3697, plain, (spl144_403 <=> (h2(e13) = h1(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl144_403])])).
fof(f4868, plain, (~ spl144_66 | ~ spl144_72 | ~ spl144_290), inference(avatar_contradiction_clause, [], [f4867])).
fof(f4867, plain, ($false | (~ spl144_66 | ~ spl144_72 | ~ spl144_290)), inference(subsumption_resolution, [], [f4866, f513])).
fof(f4866, plain, ((e21 = e22) | (~ spl144_66 | ~ spl144_72 | ~ spl144_290)), inference(forward_demodulation, [], [f4865, f1505])).
fof(f4865, plain, ((e22 = op2(e23, e23)) | (~ spl144_72 | ~ spl144_290)), inference(backward_demodulation, [], [f3010, f1530])).
fof(f3010, plain, ((e22 = op2(e23, op2(e23, e22))) | ~ spl144_290), inference(avatar_component_clause, [], [f3008])).
fof(f3008, plain, (spl144_290 <=> (e22 = op2(e23, op2(e23, e22)))), introduced(avatar_definition, [new_symbols(naming, [spl144_290])])).
fof(f4863, plain, (~ spl144_70 | ~ spl144_298), inference(avatar_split_clause, [], [f4638, f3089, f1520])).
fof(f1520, plain, (spl144_70 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_70])])).
fof(f4638, plain, (~ (e21 = op2(e23, e22)) | ~ spl144_298), inference(forward_demodulation, [], [f3061, f3090])).
fof(f3061, plain, ~ (op2(e23, e22) = h4(e13)), inference(backward_demodulation, [], [f503, f1129])).
fof(f503, plain, ~ (op2(e23, e22) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f4862, plain, (~ spl144_71 | ~ spl144_79), inference(avatar_split_clause, [], [f4861, f1558, f1524])).
fof(f1524, plain, (spl144_71 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_71])])).
fof(f1558, plain, (spl144_79 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_79])])).
fof(f4861, plain, (~ (e22 = op2(e23, e22)) | ~ spl144_79), inference(forward_demodulation, [], [f499, f1560])).
fof(f1560, plain, ((e22 = op2(e23, e20)) | ~ spl144_79), inference(avatar_component_clause, [], [f1558])).
fof(f499, plain, ~ (op2(e23, e20) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f4853, plain, (~ spl144_93 | ~ spl144_94), inference(avatar_contradiction_clause, [], [f4852])).
fof(f4852, plain, ($false | (~ spl144_93 | ~ spl144_94)), inference(subsumption_resolution, [], [f4851, f510])).
fof(f510, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f8])).
fof(f4851, plain, ((e20 = e21) | (~ spl144_93 | ~ spl144_94)), inference(backward_demodulation, [], [f1624, f1620])).
fof(f1624, plain, ((e21 = op2(e22, e20)) | ~ spl144_94), inference(avatar_component_clause, [], [f1622])).
fof(f1622, plain, (spl144_94 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_94])])).
fof(f4833, plain, (~ spl144_36 | ~ spl144_100 | ~ spl144_392 | spl144_401), inference(avatar_contradiction_clause, [], [f4832])).
fof(f4832, plain, ($false | (~ spl144_36 | ~ spl144_100 | ~ spl144_392 | spl144_401)), inference(subsumption_resolution, [], [f4831, f1649])).
fof(f1649, plain, ((e23 = op2(e21, e23)) | ~ spl144_100), inference(avatar_component_clause, [], [f1647])).
fof(f1647, plain, (spl144_100 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_100])])).
fof(f4831, plain, (~ (e23 = op2(e21, e23)) | (~ spl144_36 | ~ spl144_392 | spl144_401)), inference(forward_demodulation, [], [f4830, f3654])).
fof(f4830, plain, (~ (op2(e21, e23) = h1(e13)) | (~ spl144_36 | ~ spl144_392 | spl144_401)), inference(forward_demodulation, [], [f4829, f1345])).
fof(f1345, plain, ((e13 = op1(e11, e13)) | ~ spl144_36), inference(avatar_component_clause, [], [f1343])).
fof(f4829, plain, (~ (op2(e21, e23) = h1(op1(e11, e13))) | (~ spl144_392 | spl144_401)), inference(forward_demodulation, [], [f3691, f3654])).
fof(f3691, plain, (~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | spl144_401), inference(avatar_component_clause, [], [f3689])).
fof(f3689, plain, (spl144_401 <=> (h1(op1(e11, e13)) = op2(e21, h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl144_401])])).
fof(f4821, plain, (~ spl144_94 | ~ spl144_110), inference(avatar_contradiction_clause, [], [f4820])).
fof(f4820, plain, ($false | (~ spl144_94 | ~ spl144_110)), inference(subsumption_resolution, [], [f4819, f1692])).
fof(f4819, plain, (~ (e21 = op2(e21, e20)) | ~ spl144_94), inference(forward_demodulation, [], [f459, f1624])).
fof(f459, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f4796, plain, (spl144_69 | ~ spl144_79 | ~ spl144_292), inference(avatar_split_clause, [], [f4737, f3018, f1558, f1516])).
fof(f3018, plain, (spl144_292 <=> (e20 = op2(e23, op2(e23, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl144_292])])).
fof(f4737, plain, ((e20 = op2(e23, e22)) | (~ spl144_79 | ~ spl144_292)), inference(backward_demodulation, [], [f3020, f1560])).
fof(f3020, plain, ((e20 = op2(e23, op2(e23, e20))) | ~ spl144_292), inference(avatar_component_clause, [], [f3018])).
fof(f4780, plain, (~ spl144_28 | ~ spl144_92 | ~ spl144_392 | spl144_399), inference(avatar_contradiction_clause, [], [f4779])).
fof(f4779, plain, ($false | (~ spl144_28 | ~ spl144_92 | ~ spl144_392 | spl144_399)), inference(subsumption_resolution, [], [f4777, f3654])).
fof(f4777, plain, (~ (e23 = h1(e13)) | (~ spl144_28 | ~ spl144_92 | spl144_399)), inference(backward_demodulation, [], [f4648, f1311])).
fof(f1311, plain, ((e13 = op1(e12, e11)) | ~ spl144_28), inference(avatar_component_clause, [], [f1309])).
fof(f1309, plain, (spl144_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_28])])).
fof(f4648, plain, (~ (e23 = h1(op1(e12, e11))) | (~ spl144_92 | spl144_399)), inference(forward_demodulation, [], [f3683, f1615])).
fof(f3683, plain, (~ (op2(e22, e21) = h1(op1(e12, e11))) | spl144_399), inference(avatar_component_clause, [], [f3681])).
fof(f3681, plain, (spl144_399 <=> (op2(e22, e21) = h1(op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl144_399])])).
fof(f4764, plain, (~ spl144_23 | ~ spl144_55), inference(avatar_split_clause, [], [f4762, f1424, f1288])).
fof(f4762, plain, (~ (e12 = op1(e12, e12)) | ~ spl144_55), inference(backward_demodulation, [], [f421, f1426])).
fof(f421, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f4758, plain, (~ spl144_26 | ~ spl144_58), inference(avatar_split_clause, [], [f4752, f1437, f1301])).
fof(f4752, plain, (~ (e11 = op1(e12, e11)) | ~ spl144_58), inference(backward_demodulation, [], [f415, f1439])).
fof(f415, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f4746, plain, (~ spl144_69 | ~ spl144_72), inference(avatar_contradiction_clause, [], [f4745])).
fof(f4745, plain, ($false | (~ spl144_69 | ~ spl144_72)), inference(subsumption_resolution, [], [f4744, f512])).
fof(f4744, plain, ((e20 = e23) | (~ spl144_69 | ~ spl144_72)), inference(forward_demodulation, [], [f1530, f1518])).
fof(f4734, plain, (~ spl144_81 | ~ spl144_83), inference(avatar_contradiction_clause, [], [f4733])).
fof(f4733, plain, ($false | (~ spl144_81 | ~ spl144_83)), inference(subsumption_resolution, [], [f4732, f511])).
fof(f511, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f4732, plain, ((e20 = e22) | (~ spl144_81 | ~ spl144_83)), inference(backward_demodulation, [], [f1577, f1569])).
fof(f1569, plain, ((e20 = op2(e22, e23)) | ~ spl144_81), inference(avatar_component_clause, [], [f1567])).
fof(f1567, plain, (spl144_81 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_81])])).
fof(f4730, plain, (~ spl144_83 | ~ spl144_87 | ~ spl144_212), inference(avatar_contradiction_clause, [], [f4729])).
fof(f4729, plain, ($false | (~ spl144_83 | ~ spl144_87 | ~ spl144_212)), inference(subsumption_resolution, [], [f4728, f515])).
fof(f4728, plain, ((e22 = e23) | (~ spl144_83 | ~ spl144_87 | ~ spl144_212)), inference(forward_demodulation, [], [f4725, f1594])).
fof(f1594, plain, ((e22 = op2(e22, e22)) | ~ spl144_87), inference(avatar_component_clause, [], [f1592])).
fof(f1592, plain, (spl144_87 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_87])])).
fof(f4725, plain, ((e23 = op2(e22, e22)) | (~ spl144_83 | ~ spl144_212)), inference(backward_demodulation, [], [f2424, f1577])).
fof(f4719, plain, (spl144_89 | ~ spl144_94 | ~ spl144_215), inference(avatar_contradiction_clause, [], [f4718])).
fof(f4718, plain, ($false | (spl144_89 | ~ spl144_94 | ~ spl144_215)), inference(subsumption_resolution, [], [f4715, f1602])).
fof(f1602, plain, (~ (e20 = op2(e22, e21)) | spl144_89), inference(avatar_component_clause, [], [f1601])).
fof(f1601, plain, (spl144_89 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_89])])).
fof(f4715, plain, ((e20 = op2(e22, e21)) | (~ spl144_94 | ~ spl144_215)), inference(backward_demodulation, [], [f2439, f1624])).
fof(f2439, plain, ((e20 = op2(e22, op2(e22, e20))) | ~ spl144_215), inference(avatar_component_clause, [], [f2437])).
fof(f2437, plain, (spl144_215 <=> (e20 = op2(e22, op2(e22, e20)))), introduced(avatar_definition, [new_symbols(naming, [spl144_215])])).
fof(f4709, plain, (~ spl144_109 | ~ spl144_110), inference(avatar_contradiction_clause, [], [f4708])).
fof(f4708, plain, ($false | (~ spl144_109 | ~ spl144_110)), inference(subsumption_resolution, [], [f4707, f510])).
fof(f4707, plain, ((e20 = e21) | (~ spl144_109 | ~ spl144_110)), inference(forward_demodulation, [], [f1692, f1688])).
fof(f1688, plain, ((e20 = op2(e21, e20)) | ~ spl144_109), inference(avatar_component_clause, [], [f1686])).
fof(f1686, plain, (spl144_109 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_109])])).
fof(f4701, plain, (~ spl144_118 | ~ spl144_122), inference(avatar_split_clause, [], [f4699, f1741, f1724])).
fof(f4699, plain, (~ (e21 = op2(e20, e22)) | ~ spl144_122), inference(backward_demodulation, [], [f483, f1743])).
fof(f483, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f4653, plain, (spl144_82 | ~ spl144_92 | ~ spl144_214), inference(avatar_split_clause, [], [f4652, f2432, f1613, f1571])).
fof(f1571, plain, (spl144_82 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_82])])).
fof(f4652, plain, ((e21 = op2(e22, e23)) | (~ spl144_92 | ~ spl144_214)), inference(forward_demodulation, [], [f2434, f1615])).
fof(f2434, plain, ((e21 = op2(e22, op2(e22, e21))) | ~ spl144_214), inference(avatar_component_clause, [], [f2432])).
fof(f4646, plain, (~ spl144_81 | ~ spl144_113), inference(avatar_split_clause, [], [f4645, f1703, f1567])).
fof(f4645, plain, (~ (e20 = op2(e22, e23)) | ~ spl144_113), inference(forward_demodulation, [], [f475, f1705])).
fof(f475, plain, ~ (op2(e20, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f4644, plain, (~ spl144_82 | ~ spl144_298), inference(avatar_split_clause, [], [f4643, f3089, f1571])).
fof(f4643, plain, (~ (e21 = op2(e22, e23)) | ~ spl144_298), inference(forward_demodulation, [], [f3058, f3090])).
fof(f3058, plain, ~ (op2(e22, e23) = h4(e13)), inference(backward_demodulation, [], [f479, f1129])).
fof(f479, plain, ~ (op2(e22, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f4642, plain, (~ spl144_80 | ~ spl144_392), inference(avatar_split_clause, [], [f4280, f3653, f1562])).
fof(f1562, plain, (spl144_80 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_80])])).
fof(f4280, plain, (~ (e23 = op2(e23, e20)) | ~ spl144_392), inference(forward_demodulation, [], [f3030, f3654])).
fof(f3030, plain, ~ (op2(e23, e20) = h1(e13)), inference(backward_demodulation, [], [f458, f1117])).
fof(f458, plain, ~ (op2(e20, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f4585, plain, (~ spl144_33 | ~ spl144_49), inference(avatar_split_clause, [], [f4579, f1399, f1331])).
fof(f1331, plain, (spl144_33 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_33])])).
fof(f4579, plain, (~ (e10 = op1(e11, e13)) | ~ spl144_49), inference(backward_demodulation, [], [f426, f1401])).
fof(f426, plain, ~ (op1(e10, e13) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f4575, plain, (~ spl144_63 | ~ spl144_64), inference(avatar_contradiction_clause, [], [f4574])).
fof(f4574, plain, ($false | (~ spl144_63 | ~ spl144_64)), inference(subsumption_resolution, [], [f4573, f509])).
fof(f509, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f7])).
fof(f4573, plain, ((e12 = e13) | (~ spl144_63 | ~ spl144_64)), inference(backward_demodulation, [], [f1464, f1460])).
fof(f1460, plain, ((op1(e10, e10) = e12) | ~ spl144_63), inference(avatar_component_clause, [], [f1458])).
fof(f1458, plain, (spl144_63 <=> (op1(e10, e10) = e12)), introduced(avatar_definition, [new_symbols(naming, [spl144_63])])).
fof(f4562, plain, (~ spl144_91 | ~ spl144_92), inference(avatar_contradiction_clause, [], [f4561])).
fof(f4561, plain, ($false | (~ spl144_91 | ~ spl144_92)), inference(subsumption_resolution, [], [f4560, f515])).
fof(f4560, plain, ((e22 = e23) | (~ spl144_91 | ~ spl144_92)), inference(forward_demodulation, [], [f1615, f1611])).
fof(f1611, plain, ((e22 = op2(e22, e21)) | ~ spl144_91), inference(avatar_component_clause, [], [f1609])).
fof(f1609, plain, (spl144_91 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_91])])).
fof(f4541, plain, (~ spl144_126 | ~ spl144_128), inference(avatar_contradiction_clause, [], [f4540])).
fof(f4540, plain, ($false | (~ spl144_126 | ~ spl144_128)), inference(subsumption_resolution, [], [f4539, f514])).
fof(f4539, plain, ((e21 = e23) | (~ spl144_126 | ~ spl144_128)), inference(backward_demodulation, [], [f1768, f1760])).
fof(f1760, plain, ((op2(e20, e20) = e21) | ~ spl144_126), inference(avatar_component_clause, [], [f1758])).
fof(f1758, plain, (spl144_126 <=> (op2(e20, e20) = e21)), introduced(avatar_definition, [new_symbols(naming, [spl144_126])])).
fof(f4535, plain, (spl144_49 | ~ spl144_64 | ~ spl144_143), inference(avatar_split_clause, [], [f4534, f1879, f1462, f1399])).
fof(f1879, plain, (spl144_143 <=> (e10 = op1(e10, op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl144_143])])).
fof(f4534, plain, ((e10 = op1(e10, e13)) | (~ spl144_64 | ~ spl144_143)), inference(forward_demodulation, [], [f1881, f1464])).
fof(f1881, plain, ((e10 = op1(e10, op1(e10, e10))) | ~ spl144_143), inference(avatar_component_clause, [], [f1879])).
fof(f4520, plain, (~ spl144_91 | ~ spl144_95), inference(avatar_contradiction_clause, [], [f4519])).
fof(f4519, plain, ($false | (~ spl144_91 | ~ spl144_95)), inference(subsumption_resolution, [], [f4518, f1628])).
fof(f1628, plain, ((e22 = op2(e22, e20)) | ~ spl144_95), inference(avatar_component_clause, [], [f1626])).
fof(f1626, plain, (spl144_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_95])])).
fof(f4518, plain, (~ (e22 = op2(e22, e20)) | ~ spl144_91), inference(forward_demodulation, [], [f492, f1611])).
fof(f492, plain, ~ (op2(e22, e20) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f4515, plain, (~ spl144_85 | ~ spl144_91 | ~ spl144_214), inference(avatar_contradiction_clause, [], [f4514])).
fof(f4514, plain, ($false | (~ spl144_85 | ~ spl144_91 | ~ spl144_214)), inference(subsumption_resolution, [], [f4513, f510])).
fof(f4513, plain, ((e20 = e21) | (~ spl144_85 | ~ spl144_91 | ~ spl144_214)), inference(forward_demodulation, [], [f4512, f1586])).
fof(f1586, plain, ((e20 = op2(e22, e22)) | ~ spl144_85), inference(avatar_component_clause, [], [f1584])).
fof(f1584, plain, (spl144_85 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_85])])).
fof(f4512, plain, ((e21 = op2(e22, e22)) | (~ spl144_91 | ~ spl144_214)), inference(forward_demodulation, [], [f2434, f1611])).
fof(f4510, plain, (~ spl144_85 | spl144_314), inference(avatar_contradiction_clause, [], [f4509])).
fof(f4509, plain, ($false | (~ spl144_85 | spl144_314)), inference(subsumption_resolution, [], [f4508, f3172])).
fof(f3172, plain, (~ (e20 = h3(e13)) | spl144_314), inference(avatar_component_clause, [], [f3170])).
fof(f4508, plain, ((e20 = h3(e13)) | ~ spl144_85), inference(forward_demodulation, [], [f1125, f1586])).
fof(f4502, plain, (~ spl144_310 | ~ spl144_118), inference(avatar_split_clause, [], [f4501, f1724, f3150])).
fof(f4501, plain, (~ (e21 = h3(e13)) | ~ spl144_118), inference(forward_demodulation, [], [f3048, f1726])).
fof(f3048, plain, ~ (op2(e20, e22) = h3(e13)), inference(backward_demodulation, [], [f469, f1125])).
fof(f469, plain, ~ (op2(e20, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f4500, plain, (~ spl144_123 | ~ spl144_91), inference(avatar_split_clause, [], [f4499, f1609, f1745])).
fof(f4499, plain, (~ (e22 = op2(e20, e21)) | ~ spl144_91), inference(forward_demodulation, [], [f463, f1611])).
fof(f463, plain, ~ (op2(e20, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f4495, plain, (~ spl144_74 | ~ spl144_298), inference(avatar_split_clause, [], [f4494, f3089, f1537])).
fof(f1537, plain, (spl144_74 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_74])])).
fof(f4494, plain, (~ (e21 = op2(e23, e21)) | ~ spl144_298), inference(forward_demodulation, [], [f3060, f3090])).
fof(f3060, plain, ~ (op2(e23, e21) = h4(e13)), inference(backward_demodulation, [], [f502, f1129])).
fof(f502, plain, ~ (op2(e23, e21) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f4477, plain, (~ spl144_105 | spl144_326), inference(avatar_contradiction_clause, [], [f4476])).
fof(f4476, plain, ($false | (~ spl144_105 | spl144_326)), inference(subsumption_resolution, [], [f4475, f3233])).
fof(f3233, plain, (~ (e20 = h2(e13)) | spl144_326), inference(avatar_component_clause, [], [f3231])).
fof(f4475, plain, ((e20 = h2(e13)) | ~ spl144_105), inference(forward_demodulation, [], [f1121, f1671])).
fof(f1671, plain, ((e20 = op2(e21, e21)) | ~ spl144_105), inference(avatar_component_clause, [], [f1669])).
fof(f1669, plain, (spl144_105 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_105])])).
fof(f4461, plain, (spl144_79 | ~ spl144_392), inference(avatar_split_clause, [], [f4086, f3653, f1558])).
fof(f4086, plain, ((e22 = op2(e23, e20)) | ~ spl144_392), inference(backward_demodulation, [], [f3035, f3654])).
fof(f4450, plain, (~ spl144_41 | ~ spl144_9), inference(avatar_split_clause, [], [f4449, f1229, f1365])).
fof(f1365, plain, (spl144_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_41])])).
fof(f4449, plain, (~ (e10 = op1(e11, e11)) | ~ spl144_9), inference(forward_demodulation, [], [f418, f1231])).
fof(f418, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f4409, plain, (~ spl144_45 | ~ spl144_46), inference(avatar_contradiction_clause, [], [f4408])).
fof(f4408, plain, ($false | (~ spl144_45 | ~ spl144_46)), inference(subsumption_resolution, [], [f4407, f504])).
fof(f504, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f7])).
fof(f4407, plain, ((e10 = e11) | (~ spl144_45 | ~ spl144_46)), inference(backward_demodulation, [], [f1388, f1384])).
fof(f4400, plain, (~ spl144_53 | ~ spl144_54), inference(avatar_contradiction_clause, [], [f4399])).
fof(f4399, plain, ($false | (~ spl144_53 | ~ spl144_54)), inference(subsumption_resolution, [], [f4398, f504])).
fof(f4398, plain, ((e10 = e11) | (~ spl144_53 | ~ spl144_54)), inference(backward_demodulation, [], [f1422, f1418])).
fof(f1418, plain, ((e10 = op1(e10, e12)) | ~ spl144_53), inference(avatar_component_clause, [], [f1416])).
fof(f1416, plain, (spl144_53 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_53])])).
fof(f1422, plain, ((e11 = op1(e10, e12)) | ~ spl144_54), inference(avatar_component_clause, [], [f1420])).
fof(f4363, plain, (~ spl144_107 | ~ spl144_108), inference(avatar_contradiction_clause, [], [f4362])).
fof(f4362, plain, ($false | (~ spl144_107 | ~ spl144_108)), inference(subsumption_resolution, [], [f4361, f515])).
fof(f4361, plain, ((e22 = e23) | (~ spl144_107 | ~ spl144_108)), inference(forward_demodulation, [], [f1683, f1679])).
fof(f1679, plain, ((e22 = op2(e21, e21)) | ~ spl144_107), inference(avatar_component_clause, [], [f1677])).
fof(f1677, plain, (spl144_107 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_107])])).
fof(f1683, plain, ((e23 = op2(e21, e21)) | ~ spl144_108), inference(avatar_component_clause, [], [f1681])).
fof(f1681, plain, (spl144_108 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_108])])).
fof(f4343, plain, (~ spl144_127 | ~ spl144_128), inference(avatar_contradiction_clause, [], [f4342])).
fof(f4342, plain, ($false | (~ spl144_127 | ~ spl144_128)), inference(subsumption_resolution, [], [f4341, f515])).
fof(f4341, plain, ((e22 = e23) | (~ spl144_127 | ~ spl144_128)), inference(backward_demodulation, [], [f1768, f1764])).
fof(f1764, plain, ((op2(e20, e20) = e22) | ~ spl144_127), inference(avatar_component_clause, [], [f1762])).
fof(f1762, plain, (spl144_127 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl144_127])])).
fof(f4336, plain, (~ spl144_298 | ~ spl144_302), inference(avatar_contradiction_clause, [], [f4335])).
fof(f4335, plain, ($false | (~ spl144_298 | ~ spl144_302)), inference(subsumption_resolution, [], [f4334, f510])).
fof(f4334, plain, ((e20 = e21) | (~ spl144_298 | ~ spl144_302)), inference(forward_demodulation, [], [f3090, f3110])).
fof(f3110, plain, ((e20 = h4(e13)) | ~ spl144_302), inference(avatar_component_clause, [], [f3109])).
fof(f3109, plain, (spl144_302 <=> (e20 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_302])])).
fof(f4317, plain, (~ spl144_124 | ~ spl144_392), inference(avatar_split_clause, [], [f4316, f3653, f1749])).
fof(f1749, plain, (spl144_124 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_124])])).
fof(f4316, plain, (~ (e23 = op2(e20, e21)) | ~ spl144_392), inference(forward_demodulation, [], [f3031, f3654])).
fof(f3031, plain, ~ (op2(e20, e21) = h1(e13)), inference(backward_demodulation, [], [f480, f1117])).
fof(f480, plain, ~ (op2(e20, e20) = op2(e20, e21)), inference(cnf_transformation, [], [f6])).
fof(f4302, plain, (~ spl144_116 | ~ spl144_392), inference(avatar_split_clause, [], [f4301, f3653, f1715])).
fof(f1715, plain, (spl144_116 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_116])])).
fof(f4301, plain, (~ (e23 = op2(e20, e23)) | ~ spl144_392), inference(forward_demodulation, [], [f3033, f3654])).
fof(f3033, plain, ~ (op2(e20, e23) = h1(e13)), inference(backward_demodulation, [], [f482, f1117])).
fof(f482, plain, ~ (op2(e20, e20) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f4296, plain, (~ spl144_111 | ~ spl144_95), inference(avatar_split_clause, [], [f4295, f1626, f1694])).
fof(f1694, plain, (spl144_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_111])])).
fof(f4295, plain, (~ (e22 = op2(e21, e20)) | ~ spl144_95), inference(forward_demodulation, [], [f459, f1628])).
fof(f4274, plain, (~ spl144_77 | ~ spl144_392), inference(avatar_contradiction_clause, [], [f4273])).
fof(f4273, plain, ($false | (~ spl144_77 | ~ spl144_392)), inference(subsumption_resolution, [], [f4272, f511])).
fof(f4272, plain, ((e20 = e22) | (~ spl144_77 | ~ spl144_392)), inference(forward_demodulation, [], [f4086, f1552])).
fof(f1552, plain, ((e20 = op2(e23, e20)) | ~ spl144_77), inference(avatar_component_clause, [], [f1550])).
fof(f1550, plain, (spl144_77 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_77])])).
fof(f4255, plain, (~ spl144_50 | ~ spl144_2), inference(avatar_split_clause, [], [f4254, f1199, f1403])).
fof(f1403, plain, (spl144_50 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_50])])).
fof(f4254, plain, (~ (e11 = op1(e10, e13)) | ~ spl144_2), inference(forward_demodulation, [], [f428, f1201])).
fof(f428, plain, ~ (op1(e10, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4245, plain, (~ spl144_39 | ~ spl144_23), inference(avatar_split_clause, [], [f4244, f1288, f1356])).
fof(f4244, plain, (~ (e12 = op1(e11, e12)) | ~ spl144_23), inference(forward_demodulation, [], [f423, f1290])).
fof(f423, plain, ~ (op1(e11, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f4241, plain, (~ spl144_31 | ~ spl144_15), inference(avatar_split_clause, [], [f4240, f1254, f1322])).
fof(f1322, plain, (spl144_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_31])])).
fof(f4240, plain, (~ (e12 = op1(e12, e10)) | ~ spl144_15), inference(forward_demodulation, [], [f413, f1256])).
fof(f413, plain, ~ (op1(e12, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f4232, plain, (~ spl144_18 | ~ spl144_2), inference(avatar_split_clause, [], [f4231, f1199, f1267])).
fof(f1267, plain, (spl144_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_18])])).
fof(f4231, plain, (~ (e11 = op1(e12, e13)) | ~ spl144_2), inference(forward_demodulation, [], [f431, f1201])).
fof(f431, plain, ~ (op1(e12, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4228, plain, (~ spl144_12 | ~ spl144_8), inference(avatar_split_clause, [], [f4227, f1224, f1241])).
fof(f1241, plain, (spl144_12 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_12])])).
fof(f4227, plain, (~ (e13 = op1(e13, e11)) | ~ spl144_8), inference(forward_demodulation, [], [f453, f1226])).
fof(f453, plain, ~ (op1(e13, e11) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f4226, plain, (~ spl144_1 | ~ spl144_2), inference(avatar_contradiction_clause, [], [f4225])).
fof(f4225, plain, ($false | (~ spl144_1 | ~ spl144_2)), inference(subsumption_resolution, [], [f4224, f504])).
fof(f4224, plain, ((e10 = e11) | (~ spl144_1 | ~ spl144_2)), inference(backward_demodulation, [], [f1201, f1197])).
fof(f1197, plain, ((e10 = op1(e13, e13)) | ~ spl144_1), inference(avatar_component_clause, [], [f1195])).
fof(f1195, plain, (spl144_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_1])])).
fof(f4213, plain, (spl144_5 | ~ spl144_15 | ~ spl144_210), inference(avatar_split_clause, [], [f4210, f2406, f1254, f1212])).
fof(f2406, plain, (spl144_210 <=> (e10 = op1(e13, op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl144_210])])).
fof(f4210, plain, ((e10 = op1(e13, e12)) | (~ spl144_15 | ~ spl144_210)), inference(backward_demodulation, [], [f2408, f1256])).
fof(f2408, plain, ((e10 = op1(e13, op1(e13, e10))) | ~ spl144_210), inference(avatar_component_clause, [], [f2406])).
fof(f4212, plain, (~ spl144_7 | ~ spl144_15), inference(avatar_split_clause, [], [f4208, f1254, f1220])).
fof(f1220, plain, (spl144_7 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_7])])).
fof(f4208, plain, (~ (e12 = op1(e13, e12)) | ~ spl144_15), inference(backward_demodulation, [], [f451, f1256])).
fof(f451, plain, ~ (op1(e13, e10) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f4211, plain, (~ spl144_11 | ~ spl144_15), inference(avatar_split_clause, [], [f4207, f1254, f1237])).
fof(f1237, plain, (spl144_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_11])])).
fof(f4207, plain, (~ (e12 = op1(e13, e11)) | ~ spl144_15), inference(backward_demodulation, [], [f450, f1256])).
fof(f450, plain, ~ (op1(e13, e10) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f4206, plain, (~ spl144_22 | ~ spl144_23), inference(avatar_contradiction_clause, [], [f4205])).
fof(f4205, plain, ($false | (~ spl144_22 | ~ spl144_23)), inference(subsumption_resolution, [], [f4204, f507])).
fof(f4204, plain, ((e11 = e12) | (~ spl144_22 | ~ spl144_23)), inference(backward_demodulation, [], [f1290, f1286])).
fof(f4199, plain, (~ spl144_12 | ~ spl144_28), inference(avatar_split_clause, [], [f4196, f1309, f1241])).
fof(f4196, plain, (~ (e13 = op1(e13, e11)) | ~ spl144_28), inference(backward_demodulation, [], [f419, f1311])).
fof(f4188, plain, (~ spl144_12 | ~ spl144_44), inference(avatar_split_clause, [], [f4185, f1377, f1241])).
fof(f4185, plain, (~ (e13 = op1(e13, e11)) | ~ spl144_44), inference(backward_demodulation, [], [f418, f1379])).
fof(f4172, plain, (~ spl144_57 | ~ spl144_59), inference(avatar_contradiction_clause, [], [f4171])).
fof(f4171, plain, ($false | (~ spl144_57 | ~ spl144_59)), inference(subsumption_resolution, [], [f4170, f505])).
fof(f505, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f7])).
fof(f4170, plain, ((e10 = e12) | (~ spl144_57 | ~ spl144_59)), inference(forward_demodulation, [], [f1443, f1435])).
fof(f1435, plain, ((e10 = op1(e10, e11)) | ~ spl144_57), inference(avatar_component_clause, [], [f1433])).
fof(f1433, plain, (spl144_57 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_57])])).
fof(f4168, plain, (spl144_298 | ~ spl144_66), inference(avatar_split_clause, [], [f4167, f1503, f3089])).
fof(f4167, plain, ((e21 = h4(e13)) | ~ spl144_66), inference(backward_demodulation, [], [f1129, f1505])).
fof(f4160, plain, (~ spl144_81 | ~ spl144_84), inference(avatar_contradiction_clause, [], [f4159])).
fof(f4159, plain, ($false | (~ spl144_81 | ~ spl144_84)), inference(subsumption_resolution, [], [f4158, f512])).
fof(f4158, plain, ((e20 = e23) | (~ spl144_81 | ~ spl144_84)), inference(backward_demodulation, [], [f1581, f1569])).
fof(f4156, plain, (~ spl144_362 | ~ spl144_84), inference(avatar_split_clause, [], [f4153, f1579, f3449])).
fof(f4153, plain, (~ (e23 = h3(e13)) | ~ spl144_84), inference(backward_demodulation, [], [f3053, f1581])).
fof(f3053, plain, ~ (op2(e22, e23) = h3(e13)), inference(backward_demodulation, [], [f497, f1125])).
fof(f497, plain, ~ (op2(e22, e22) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f4146, plain, (~ spl144_83 | ~ spl144_95), inference(avatar_split_clause, [], [f4142, f1626, f1575])).
fof(f4142, plain, (~ (e22 = op2(e22, e23)) | ~ spl144_95), inference(backward_demodulation, [], [f494, f1628])).
fof(f494, plain, ~ (op2(e22, e20) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f4134, plain, (~ spl144_106 | ~ spl144_107), inference(avatar_contradiction_clause, [], [f4133])).
fof(f4133, plain, ($false | (~ spl144_106 | ~ spl144_107)), inference(subsumption_resolution, [], [f4132, f513])).
fof(f4132, plain, ((e21 = e22) | (~ spl144_106 | ~ spl144_107)), inference(backward_demodulation, [], [f1679, f1675])).
fof(f1675, plain, ((e21 = op2(e21, e21)) | ~ spl144_106), inference(avatar_component_clause, [], [f1673])).
fof(f1673, plain, (spl144_106 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_106])])).
fof(f4116, plain, (spl144_120 | ~ spl144_115 | ~ spl144_222), inference(avatar_split_clause, [], [f4115, f2476, f1711, f1732])).
fof(f1732, plain, (spl144_120 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_120])])).
fof(f2476, plain, (spl144_222 <=> (e23 = op2(e20, op2(e20, e23)))), introduced(avatar_definition, [new_symbols(naming, [spl144_222])])).
fof(f4115, plain, ((e23 = op2(e20, e22)) | (~ spl144_115 | ~ spl144_222)), inference(forward_demodulation, [], [f2478, f1713])).
fof(f1713, plain, ((e22 = op2(e20, e23)) | ~ spl144_115), inference(avatar_component_clause, [], [f1711])).
fof(f2478, plain, ((e23 = op2(e20, op2(e20, e23))) | ~ spl144_222), inference(avatar_component_clause, [], [f2476])).
fof(f4102, plain, (~ spl144_103 | ~ spl144_318), inference(avatar_split_clause, [], [f4096, f3190, f1660])).
fof(f4096, plain, (~ (e22 = op2(e21, e22)) | ~ spl144_318), inference(backward_demodulation, [], [f3044, f3191])).
fof(f3044, plain, ~ (op2(e21, e22) = h2(e13)), inference(backward_demodulation, [], [f489, f1121])).
fof(f489, plain, ~ (op2(e21, e21) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f4093, plain, (~ spl144_78 | ~ spl144_392), inference(avatar_contradiction_clause, [], [f4092])).
fof(f4092, plain, ($false | (~ spl144_78 | ~ spl144_392)), inference(subsumption_resolution, [], [f4091, f513])).
fof(f4091, plain, ((e21 = e22) | (~ spl144_78 | ~ spl144_392)), inference(forward_demodulation, [], [f4086, f1556])).
fof(f1556, plain, ((e21 = op2(e23, e20)) | ~ spl144_78), inference(avatar_component_clause, [], [f1554])).
fof(f1554, plain, (spl144_78 <=> (e21 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_78])])).
fof(f4090, plain, (spl144_66 | ~ spl144_392), inference(avatar_split_clause, [], [f4085, f3653, f1503])).
fof(f4085, plain, ((e21 = op2(e23, e23)) | ~ spl144_392), inference(backward_demodulation, [], [f3034, f3654])).
fof(f4089, plain, (~ spl144_120 | ~ spl144_392), inference(avatar_split_clause, [], [f4084, f3653, f1732])).
fof(f4084, plain, (~ (e23 = op2(e20, e22)) | ~ spl144_392), inference(backward_demodulation, [], [f3032, f3654])).
fof(f3032, plain, ~ (op2(e20, e22) = h1(e13)), inference(backward_demodulation, [], [f481, f1117])).
fof(f481, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f4088, plain, (~ spl144_96 | ~ spl144_392), inference(avatar_split_clause, [], [f4083, f3653, f1630])).
fof(f1630, plain, (spl144_96 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_96])])).
fof(f4083, plain, (~ (e23 = op2(e22, e20)) | ~ spl144_392), inference(backward_demodulation, [], [f3029, f3654])).
fof(f3029, plain, ~ (op2(e22, e20) = h1(e13)), inference(backward_demodulation, [], [f457, f1117])).
fof(f457, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f4087, plain, (~ spl144_112 | ~ spl144_392), inference(avatar_split_clause, [], [f4082, f3653, f1698])).
fof(f1698, plain, (spl144_112 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl144_112])])).
fof(f4082, plain, (~ (e23 = op2(e21, e20)) | ~ spl144_392), inference(backward_demodulation, [], [f3028, f3654])).
fof(f3028, plain, ~ (op2(e21, e20) = h1(e13)), inference(backward_demodulation, [], [f456, f1117])).
fof(f456, plain, ~ (op2(e20, e20) = op2(e21, e20)), inference(cnf_transformation, [], [f6])).
fof(f4071, plain, (~ spl144_298 | ~ spl144_98), inference(avatar_split_clause, [], [f4070, f1639, f3089])).
fof(f1639, plain, (spl144_98 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_98])])).
fof(f4070, plain, (~ (e21 = h4(e13)) | ~ spl144_98), inference(forward_demodulation, [], [f3057, f1641])).
fof(f1641, plain, ((e21 = op2(e21, e23)) | ~ spl144_98), inference(avatar_component_clause, [], [f1639])).
fof(f3057, plain, ~ (op2(e21, e23) = h4(e13)), inference(backward_demodulation, [], [f478, f1129])).
fof(f478, plain, ~ (op2(e21, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f4061, plain, (spl144_310 | ~ spl144_86), inference(avatar_split_clause, [], [f4060, f1588, f3150])).
fof(f1588, plain, (spl144_86 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_86])])).
fof(f4060, plain, ((e21 = h3(e13)) | ~ spl144_86), inference(forward_demodulation, [], [f1125, f1590])).
fof(f1590, plain, ((e21 = op2(e22, e22)) | ~ spl144_86), inference(avatar_component_clause, [], [f1588])).
fof(f4025, plain, (~ spl144_34 | ~ spl144_2), inference(avatar_split_clause, [], [f4024, f1199, f1335])).
fof(f1335, plain, (spl144_34 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_34])])).
fof(f4024, plain, (~ (e11 = op1(e11, e13)) | ~ spl144_2), inference(forward_demodulation, [], [f430, f1201])).
fof(f430, plain, ~ (op1(e11, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4019, plain, (~ spl144_14 | ~ spl144_2), inference(avatar_split_clause, [], [f4018, f1199, f1250])).
fof(f1250, plain, (spl144_14 <=> (e11 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_14])])).
fof(f4018, plain, (~ (e11 = op1(e13, e10)) | ~ spl144_2), inference(forward_demodulation, [], [f452, f1201])).
fof(f452, plain, ~ (op1(e13, e10) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4017, plain, (~ spl144_10 | ~ spl144_2), inference(avatar_split_clause, [], [f4016, f1199, f1233])).
fof(f1233, plain, (spl144_10 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_10])])).
fof(f4016, plain, (~ (e11 = op1(e13, e11)) | ~ spl144_2), inference(forward_demodulation, [], [f454, f1201])).
fof(f454, plain, ~ (op1(e13, e11) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4014, plain, (~ spl144_6 | ~ spl144_2), inference(avatar_split_clause, [], [f4013, f1199, f1216])).
fof(f1216, plain, (spl144_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_6])])).
fof(f4013, plain, (~ (e11 = op1(e13, e12)) | ~ spl144_2), inference(forward_demodulation, [], [f455, f1201])).
fof(f455, plain, ~ (op1(e13, e12) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f3984, plain, (~ spl144_19 | ~ spl144_20), inference(avatar_contradiction_clause, [], [f3983])).
fof(f3983, plain, ($false | (~ spl144_19 | ~ spl144_20)), inference(subsumption_resolution, [], [f3982, f509])).
fof(f3982, plain, ((e12 = e13) | (~ spl144_19 | ~ spl144_20)), inference(backward_demodulation, [], [f1277, f1273])).
fof(f1277, plain, ((e13 = op1(e12, e13)) | ~ spl144_20), inference(avatar_component_clause, [], [f1275])).
fof(f1275, plain, (spl144_20 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_20])])).
fof(f3978, plain, (~ spl144_5 | ~ spl144_21), inference(avatar_split_clause, [], [f3976, f1280, f1212])).
fof(f1280, plain, (spl144_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_21])])).
fof(f3976, plain, (~ (e10 = op1(e13, e12)) | ~ spl144_21), inference(backward_demodulation, [], [f425, f1282])).
fof(f1282, plain, ((e10 = op1(e12, e12)) | ~ spl144_21), inference(avatar_component_clause, [], [f1280])).
fof(f3965, plain, (~ spl144_21 | ~ spl144_29), inference(avatar_split_clause, [], [f3961, f1314, f1280])).
fof(f3961, plain, (~ (e10 = op1(e12, e12)) | ~ spl144_29), inference(backward_demodulation, [], [f445, f1316])).
fof(f3957, plain, (~ spl144_3 | ~ spl144_35), inference(avatar_split_clause, [], [f3954, f1339, f1203])).
fof(f1203, plain, (spl144_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_3])])).
fof(f3954, plain, (~ (e12 = op1(e13, e13)) | ~ spl144_35), inference(backward_demodulation, [], [f430, f1341])).
fof(f3949, plain, (~ spl144_38 | ~ spl144_39), inference(avatar_contradiction_clause, [], [f3948])).
fof(f3948, plain, ($false | (~ spl144_38 | ~ spl144_39)), inference(subsumption_resolution, [], [f3947, f507])).
fof(f3947, plain, ((e11 = e12) | (~ spl144_38 | ~ spl144_39)), inference(backward_demodulation, [], [f1358, f1354])).
fof(f1354, plain, ((e11 = op1(e11, e12)) | ~ spl144_38), inference(avatar_component_clause, [], [f1352])).
fof(f1358, plain, ((e12 = op1(e11, e12)) | ~ spl144_39), inference(avatar_component_clause, [], [f1356])).
fof(f3934, plain, (~ spl144_27 | ~ spl144_43), inference(avatar_split_clause, [], [f3930, f1373, f1305])).
fof(f3930, plain, (~ (e12 = op1(e12, e11)) | ~ spl144_43), inference(backward_demodulation, [], [f417, f1375])).
fof(f417, plain, ~ (op1(e11, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f3922, plain, (~ spl144_15 | ~ spl144_47), inference(avatar_split_clause, [], [f3917, f1390, f1254])).
fof(f1390, plain, (spl144_47 <=> (e12 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_47])])).
fof(f3917, plain, (~ (e12 = op1(e13, e10)) | ~ spl144_47), inference(backward_demodulation, [], [f412, f1392])).
fof(f1392, plain, ((e12 = op1(e11, e10)) | ~ spl144_47), inference(avatar_component_clause, [], [f1390])).
fof(f412, plain, ~ (op1(e11, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f3909, plain, (~ spl144_3 | ~ spl144_51), inference(avatar_split_clause, [], [f3906, f1407, f1203])).
fof(f1407, plain, (spl144_51 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_51])])).
fof(f3906, plain, (~ (e12 = op1(e13, e13)) | ~ spl144_51), inference(backward_demodulation, [], [f428, f1409])).
fof(f1409, plain, ((e12 = op1(e10, e13)) | ~ spl144_51), inference(avatar_component_clause, [], [f1407])).
fof(f3908, plain, (~ spl144_19 | ~ spl144_51), inference(avatar_split_clause, [], [f3905, f1407, f1271])).
fof(f3905, plain, (~ (e12 = op1(e12, e13)) | ~ spl144_51), inference(backward_demodulation, [], [f427, f1409])).
fof(f427, plain, ~ (op1(e10, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f3901, plain, (~ spl144_49 | ~ spl144_53), inference(avatar_split_clause, [], [f3896, f1416, f1399])).
fof(f3896, plain, (~ (e10 = op1(e10, e13)) | ~ spl144_53), inference(backward_demodulation, [], [f437, f1418])).
fof(f437, plain, ~ (op1(e10, e12) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f3900, plain, (~ spl144_5 | ~ spl144_53), inference(avatar_split_clause, [], [f3895, f1416, f1212])).
fof(f3895, plain, (~ (e10 = op1(e13, e12)) | ~ spl144_53), inference(backward_demodulation, [], [f422, f1418])).
fof(f422, plain, ~ (op1(e10, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f3891, plain, (~ spl144_49 | ~ spl144_57), inference(avatar_split_clause, [], [f3885, f1433, f1399])).
fof(f3885, plain, (~ (e10 = op1(e10, e13)) | ~ spl144_57), inference(backward_demodulation, [], [f436, f1435])).
fof(f436, plain, ~ (op1(e10, e11) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f3889, plain, (~ spl144_9 | ~ spl144_57), inference(avatar_split_clause, [], [f3883, f1433, f1229])).
fof(f3883, plain, (~ (e10 = op1(e13, e11)) | ~ spl144_57), inference(backward_demodulation, [], [f416, f1435])).
fof(f416, plain, ~ (op1(e10, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f3880, plain, (spl144_15 | ~ spl144_64), inference(avatar_split_clause, [], [f3869, f1462, f1254])).
fof(f3869, plain, ((e12 = op1(e13, e10)) | ~ spl144_64), inference(backward_demodulation, [], [f1109, f1464])).
fof(f1109, plain, (e12 = op1(op1(e10, e10), e10)), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((op1(e10, e10) = e13) & (e12 = op1(op1(e10, e10), e10)) & (e11 = op1(op1(e10, e10), op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax12)).
fof(f3879, plain, (spl144_2 | ~ spl144_64), inference(avatar_split_clause, [], [f3868, f1462, f1199])).
fof(f3868, plain, ((e11 = op1(e13, e13)) | ~ spl144_64), inference(backward_demodulation, [], [f1108, f1464])).
fof(f1108, plain, (e11 = op1(op1(e10, e10), op1(e10, e10))), inference(cnf_transformation, [], [f12])).
fof(f3878, plain, (~ spl144_52 | ~ spl144_64), inference(avatar_split_clause, [], [f3867, f1462, f1411])).
fof(f1411, plain, (spl144_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl144_52])])).
fof(f3867, plain, (~ (e13 = op1(e10, e13)) | ~ spl144_64), inference(backward_demodulation, [], [f434, f1464])).
fof(f434, plain, ~ (op1(e10, e10) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f3877, plain, (~ spl144_56 | ~ spl144_64), inference(avatar_split_clause, [], [f3866, f1462, f1428])).
fof(f1428, plain, (spl144_56 <=> (e13 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl144_56])])).
fof(f3866, plain, (~ (e13 = op1(e10, e12)) | ~ spl144_64), inference(backward_demodulation, [], [f433, f1464])).
fof(f433, plain, ~ (op1(e10, e10) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f3876, plain, (~ spl144_60 | ~ spl144_64), inference(avatar_split_clause, [], [f3865, f1462, f1445])).
fof(f1445, plain, (spl144_60 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl144_60])])).
fof(f3865, plain, (~ (e13 = op1(e10, e11)) | ~ spl144_64), inference(backward_demodulation, [], [f432, f1464])).
fof(f432, plain, ~ (op1(e10, e10) = op1(e10, e11)), inference(cnf_transformation, [], [f5])).
fof(f3875, plain, (~ spl144_16 | ~ spl144_64), inference(avatar_split_clause, [], [f3864, f1462, f1258])).
fof(f1258, plain, (spl144_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_16])])).
fof(f3864, plain, (~ (e13 = op1(e13, e10)) | ~ spl144_64), inference(backward_demodulation, [], [f410, f1464])).
fof(f410, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f3874, plain, (~ spl144_32 | ~ spl144_64), inference(avatar_split_clause, [], [f3863, f1462, f1326])).
fof(f1326, plain, (spl144_32 <=> (e13 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_32])])).
fof(f3863, plain, (~ (e13 = op1(e12, e10)) | ~ spl144_64), inference(backward_demodulation, [], [f409, f1464])).
fof(f409, plain, ~ (op1(e10, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f3873, plain, (~ spl144_48 | ~ spl144_64), inference(avatar_split_clause, [], [f3862, f1462, f1394])).
fof(f1394, plain, (spl144_48 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_48])])).
fof(f3862, plain, (~ (e13 = op1(e11, e10)) | ~ spl144_64), inference(backward_demodulation, [], [f408, f1464])).
fof(f408, plain, ~ (op1(e10, e10) = op1(e11, e10)), inference(cnf_transformation, [], [f5])).
fof(f3861, plain, (spl144_302 | ~ spl144_65), inference(avatar_split_clause, [], [f3860, f1499, f3109])).
fof(f1499, plain, (spl144_65 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl144_65])])).
fof(f3860, plain, ((e20 = h4(e13)) | ~ spl144_65), inference(backward_demodulation, [], [f1129, f1501])).
fof(f1501, plain, ((e20 = op2(e23, e23)) | ~ spl144_65), inference(avatar_component_clause, [], [f1499])).
fof(f3838, plain, (~ spl144_75 | ~ spl144_79), inference(avatar_split_clause, [], [f3834, f1558, f1541])).
fof(f1541, plain, (spl144_75 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_75])])).
fof(f3834, plain, (~ (e22 = op2(e23, e21)) | ~ spl144_79), inference(backward_demodulation, [], [f498, f1560])).
fof(f498, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f3818, plain, (spl144_362 | ~ spl144_88), inference(avatar_split_clause, [], [f3817, f1596, f3449])).
fof(f3817, plain, ((e23 = h3(e13)) | ~ spl144_88), inference(backward_demodulation, [], [f1125, f1598])).
fof(f1598, plain, ((e23 = op2(e22, e22)) | ~ spl144_88), inference(avatar_component_clause, [], [f1596])).
fof(f3816, plain, (~ spl144_314 | ~ spl144_89), inference(avatar_split_clause, [], [f3811, f1601, f3170])).
fof(f3811, plain, (~ (e20 = h3(e13)) | ~ spl144_89), inference(backward_demodulation, [], [f3052, f1603])).
fof(f1603, plain, ((e20 = op2(e22, e21)) | ~ spl144_89), inference(avatar_component_clause, [], [f1601])).
fof(f3812, plain, (~ spl144_73 | ~ spl144_89), inference(avatar_split_clause, [], [f3806, f1601, f1533])).
fof(f3806, plain, (~ (e20 = op2(e23, e21)) | ~ spl144_89), inference(backward_demodulation, [], [f467, f1603])).
fof(f3781, plain, (~ spl144_310 | ~ spl144_102), inference(avatar_split_clause, [], [f3777, f1656, f3150])).
fof(f3777, plain, (~ (e21 = h3(e13)) | ~ spl144_102), inference(backward_demodulation, [], [f3049, f1658])).
fof(f3049, plain, ~ (op2(e21, e22) = h3(e13)), inference(backward_demodulation, [], [f471, f1125])).
fof(f471, plain, ~ (op2(e21, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f3776, plain, (spl144_318 | ~ spl144_107), inference(avatar_split_clause, [], [f3775, f1677, f3190])).
fof(f3775, plain, ((e22 = h2(e13)) | ~ spl144_107), inference(backward_demodulation, [], [f1121, f1679])).
fof(f3773, plain, (~ spl144_326 | ~ spl144_109), inference(avatar_split_clause, [], [f3767, f1686, f3231])).
fof(f3767, plain, (~ (e20 = h2(e13)) | ~ spl144_109), inference(backward_demodulation, [], [f3043, f1688])).
fof(f3043, plain, ~ (op2(e21, e20) = h2(e13)), inference(backward_demodulation, [], [f486, f1121])).
fof(f486, plain, ~ (op2(e21, e20) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f3770, plain, (~ spl144_101 | ~ spl144_109), inference(avatar_split_clause, [], [f3763, f1686, f1652])).
fof(f3763, plain, (~ (e20 = op2(e21, e22)) | ~ spl144_109), inference(backward_demodulation, [], [f487, f1688])).
fof(f3749, plain, (~ spl144_314 | ~ spl144_117), inference(avatar_split_clause, [], [f3743, f1720, f3170])).
fof(f1720, plain, (spl144_117 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl144_117])])).
fof(f3743, plain, (~ (e20 = h3(e13)) | ~ spl144_117), inference(backward_demodulation, [], [f3048, f1722])).
fof(f1722, plain, ((e20 = op2(e20, e22)) | ~ spl144_117), inference(avatar_component_clause, [], [f1720])).
fof(f3747, plain, (~ spl144_113 | ~ spl144_117), inference(avatar_split_clause, [], [f3741, f1720, f1703])).
fof(f3741, plain, (~ (e20 = op2(e20, e23)) | ~ spl144_117), inference(backward_demodulation, [], [f485, f1722])).
fof(f3746, plain, (~ spl144_69 | ~ spl144_117), inference(avatar_split_clause, [], [f3740, f1720, f1516])).
fof(f3740, plain, (~ (e20 = op2(e23, e22)) | ~ spl144_117), inference(backward_demodulation, [], [f470, f1722])).
fof(f470, plain, ~ (op2(e20, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f3745, plain, (~ spl144_101 | ~ spl144_117), inference(avatar_split_clause, [], [f3739, f1720, f1652])).
fof(f3739, plain, (~ (e20 = op2(e21, e22)) | ~ spl144_117), inference(backward_demodulation, [], [f468, f1722])).
fof(f468, plain, ~ (op2(e20, e22) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f3736, plain, (~ spl144_73 | ~ spl144_121), inference(avatar_split_clause, [], [f3730, f1737, f1533])).
fof(f1737, plain, (spl144_121 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl144_121])])).
fof(f3730, plain, (~ (e20 = op2(e23, e21)) | ~ spl144_121), inference(backward_demodulation, [], [f464, f1739])).
fof(f1739, plain, ((e20 = op2(e20, e21)) | ~ spl144_121), inference(avatar_component_clause, [], [f1737])).
fof(f464, plain, ~ (op2(e20, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f3735, plain, (~ spl144_89 | ~ spl144_121), inference(avatar_split_clause, [], [f3729, f1737, f1601])).
fof(f3729, plain, (~ (e20 = op2(e22, e21)) | ~ spl144_121), inference(backward_demodulation, [], [f463, f1739])).
fof(f3726, plain, (spl144_392 | ~ spl144_128), inference(avatar_split_clause, [], [f3725, f1766, f3653])).
fof(f3725, plain, ((e23 = h1(e13)) | ~ spl144_128), inference(backward_demodulation, [], [f1117, f1768])).
fof(f3720, plain, (spl144_333 | spl144_331 | spl144_329 | ~ spl144_392 | ~ spl144_393 | ~ spl144_394 | ~ spl144_395 | ~ spl144_396 | ~ spl144_397 | ~ spl144_398 | ~ spl144_399 | ~ spl144_400 | ~ spl144_401 | ~ spl144_402 | ~ spl144_403 | ~ spl144_404 | ~ spl144_405 | ~ spl144_406 | ~ spl144_407 | ~ spl144_408), inference(avatar_split_clause, [], [f3651, f3717, f3713, f3709, f3705, f3701, f3697, f3693, f3689, f3685, f3681, f3677, f3673, f3669, f3665, f3661, f3657, f3653, f3247, f3260, f3273])).
fof(f3273, plain, (spl144_333 <=> sP132), introduced(avatar_definition, [new_symbols(naming, [spl144_333])])).
fof(f3260, plain, (spl144_331 <=> sP133), introduced(avatar_definition, [new_symbols(naming, [spl144_331])])).
fof(f3247, plain, (spl144_329 <=> sP134), introduced(avatar_definition, [new_symbols(naming, [spl144_329])])).
fof(f3651, plain, (~ (h1(e13) = h1(op1(e10, e10))) | ~ (op2(e20, e21) = h1(op1(e10, e11))) | ~ (op2(e20, e22) = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e13) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132), inference(forward_demodulation, [], [f3650, f1117])).
fof(f3650, plain, (~ (op2(e20, e20) = h1(op1(e10, e10))) | ~ (op2(e20, e21) = h1(op1(e10, e11))) | ~ (op2(e20, e22) = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e13) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132), inference(forward_demodulation, [], [f3649, f1114])).
fof(f3649, plain, (~ (op2(e20, e21) = h1(op1(e10, e11))) | ~ (op2(e20, e22) = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e13) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3648, f1114])).
fof(f3648, plain, (~ (h1(op1(e10, e11)) = op2(h1(e10), e21)) | ~ (op2(e20, e22) = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e13) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3647, f3039])).
fof(f3647, plain, (~ (op2(e20, e22) = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e13) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3646, f1114])).
fof(f3646, plain, (~ (h1(op1(e10, e12)) = op2(h1(e10), e22)) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e13) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3645, f3037])).
fof(f3645, plain, (~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e13) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3644, f1114])).
fof(f3644, plain, (~ (op2(e21, e20) = h1(op1(e11, e10))) | ~ (h2(e13) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3643, f3039])).
fof(f3643, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | ~ (h2(e13) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3642, f1114])).
fof(f3642, plain, (~ (h2(e13) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3641, f1121])).
fof(f3641, plain, (~ (op2(e21, e21) = h1(op1(e11, e11))) | ~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3640, f3039])).
fof(f3640, plain, (~ (op2(e21, e22) = h1(op1(e11, e12))) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3639, f3039])).
fof(f3639, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | ~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3638, f3037])).
fof(f3638, plain, (~ (h1(op1(e11, e13)) = op2(e21, h1(e13))) | ~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3637, f3039])).
fof(f3637, plain, (~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3636, f3037])).
fof(f3636, plain, (~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3635, f1114])).
fof(f3635, plain, (~ (op2(e22, e21) = h1(op1(e12, e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3634, f3037])).
fof(f3634, plain, (~ (h1(op1(e12, e11)) = op2(h1(e12), e21)) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3633, f3039])).
fof(f3633, plain, (~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3632, f1125])).
fof(f3632, plain, (~ (op2(e22, e22) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3631, f3037])).
fof(f3631, plain, (~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3630, f3037])).
fof(f3630, plain, (~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3629, f3035])).
fof(f3629, plain, (~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3628, f1114])).
fof(f3628, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3627, f3039])).
fof(f3627, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f3626, f3037])).
fof(f3626, plain, (~ (e21 = h1(op1(e13, e13))) | ~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1181, f3034])).
fof(f1181, plain, (~ (e23 = h1(e13)) | sP134 | sP133 | sP132 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP143 | sP142 | sP141 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP140 | sP139 | sP138 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP137 | sP136 | sP135 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP134 | sP133 | sP132 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f20, e166, e165, e164, e163, e162, e161, e160, e159, e158, e157, e156, e155])).
fof(f155, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP132), inference(usedef, [], [e155])).
fof(e155, plain, (sP132 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP132])])).
fof(f156, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP133), inference(usedef, [], [e156])).
fof(e156, plain, (sP133 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP133])])).
fof(f157, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP134), inference(usedef, [], [e157])).
fof(e157, plain, (sP134 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP134])])).
fof(f158, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP135), inference(usedef, [], [e158])).
fof(e158, plain, (sP135 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP135])])).
fof(f159, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP136), inference(usedef, [], [e159])).
fof(e159, plain, (sP136 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP136])])).
fof(f160, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP137), inference(usedef, [], [e160])).
fof(e160, plain, (sP137 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP137])])).
fof(f161, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP138), inference(usedef, [], [e161])).
fof(e161, plain, (sP138 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP138])])).
fof(f162, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP139), inference(usedef, [], [e162])).
fof(e162, plain, (sP139 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP139])])).
fof(f163, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP140), inference(usedef, [], [e163])).
fof(e163, plain, (sP140 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP140])])).
fof(f164, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP141), inference(usedef, [], [e164])).
fof(e164, plain, (sP141 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP141])])).
fof(f165, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP142), inference(usedef, [], [e165])).
fof(e165, plain, (sP142 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP142])])).
fof(f166, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP143), inference(usedef, [], [e166])).
fof(e166, plain, (sP143 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP143])])).
fof(f20, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f18])).
fof(f18, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', co1)).
fof(f3284, plain, ~ spl144_333, inference(avatar_split_clause, [], [f3283, f3273])).
fof(f3283, plain, ~ sP132, inference(subsumption_resolution, [], [f1174, f1114])).
fof(f1174, plain, (~ (e20 = h1(e10)) | ~ sP132), inference(cnf_transformation, [], [f311])).
fof(f311, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP132), inference(nnf_transformation, [], [f155])).
fof(f3270, plain, ~ spl144_331, inference(avatar_split_clause, [], [f3269, f3260])).
fof(f3269, plain, ~ sP133, inference(subsumption_resolution, [], [f1171, f3039])).
fof(f1171, plain, (~ (e21 = h1(e11)) | ~ sP133), inference(cnf_transformation, [], [f310])).
fof(f310, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP133), inference(nnf_transformation, [], [f156])).
fof(f3256, plain, ~ spl144_329, inference(avatar_split_clause, [], [f3255, f3247])).
fof(f3255, plain, ~ sP134, inference(subsumption_resolution, [], [f1168, f3037])).
fof(f1168, plain, (~ (e22 = h1(e12)) | ~ sP134), inference(cnf_transformation, [], [f309])).
fof(f309, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP134), inference(nnf_transformation, [], [f157])).
fof(f3027, plain, spl144_128, inference(avatar_split_clause, [], [f1113, f1766])).
fof(f1113, plain, (op2(e20, e20) = e23), inference(cnf_transformation, [], [f13])).
fof(f3026, plain, spl144_64, inference(avatar_split_clause, [], [f1110, f1462])).
fof(f1110, plain, (op1(e10, e10) = e13), inference(cnf_transformation, [], [f12])).
fof(f3025, plain, (spl144_288 | spl144_287 | spl144_286 | spl144_285 | spl144_284 | spl144_283 | spl144_282 | spl144_281 | spl144_280 | spl144_279 | spl144_278 | spl144_277 | spl144_276 | spl144_275 | spl144_274 | spl144_273 | spl144_272 | spl144_271 | spl144_270 | spl144_269 | spl144_268 | spl144_267 | spl144_266 | spl144_265 | spl144_264 | spl144_263 | spl144_262 | spl144_261 | spl144_260 | spl144_259 | spl144_258 | spl144_257 | spl144_256 | spl144_255 | spl144_254 | spl144_253 | spl144_252 | spl144_251 | spl144_250 | spl144_249 | spl144_248 | spl144_247 | spl144_246 | spl144_245 | spl144_244 | spl144_243 | spl144_242 | spl144_241 | spl144_240 | spl144_239 | spl144_238 | spl144_237 | spl144_236 | spl144_235 | spl144_234 | spl144_233 | spl144_232 | spl144_231 | spl144_230 | spl144_229 | spl144_228 | spl144_227 | spl144_226 | spl144_68), inference(avatar_split_clause, [], [f1096, f1511, f2496, f2504, f2512, f2520, f2528, f2536, f2544, f2552, f2560, f2568, f2576, f2584, f2592, f2600, f2608, f2616, f2624, f2632, f2640, f2648, f2656, f2664, f2672, f2680, f2688, f2696, f2704, f2712, f2720, f2728, f2736, f2744, f2752, f2760, f2768, f2776, f2784, f2792, f2800, f2808, f2816, f2824, f2832, f2840, f2848, f2856, f2864, f2872, f2880, f2888, f2896, f2904, f2912, f2920, f2928, f2936, f2944, f2952, f2960, f2968, f2976, f2984, f2992])).
fof(f2992, plain, (spl144_288 <=> sP66), introduced(avatar_definition, [new_symbols(naming, [spl144_288])])).
fof(f2984, plain, (spl144_287 <=> sP67), introduced(avatar_definition, [new_symbols(naming, [spl144_287])])).
fof(f2976, plain, (spl144_286 <=> sP68), introduced(avatar_definition, [new_symbols(naming, [spl144_286])])).
fof(f2968, plain, (spl144_285 <=> sP69), introduced(avatar_definition, [new_symbols(naming, [spl144_285])])).
fof(f2960, plain, (spl144_284 <=> sP70), introduced(avatar_definition, [new_symbols(naming, [spl144_284])])).
fof(f2952, plain, (spl144_283 <=> sP71), introduced(avatar_definition, [new_symbols(naming, [spl144_283])])).
fof(f2944, plain, (spl144_282 <=> sP72), introduced(avatar_definition, [new_symbols(naming, [spl144_282])])).
fof(f2936, plain, (spl144_281 <=> sP73), introduced(avatar_definition, [new_symbols(naming, [spl144_281])])).
fof(f2928, plain, (spl144_280 <=> sP74), introduced(avatar_definition, [new_symbols(naming, [spl144_280])])).
fof(f2920, plain, (spl144_279 <=> sP75), introduced(avatar_definition, [new_symbols(naming, [spl144_279])])).
fof(f2912, plain, (spl144_278 <=> sP76), introduced(avatar_definition, [new_symbols(naming, [spl144_278])])).
fof(f2904, plain, (spl144_277 <=> sP77), introduced(avatar_definition, [new_symbols(naming, [spl144_277])])).
fof(f2896, plain, (spl144_276 <=> sP78), introduced(avatar_definition, [new_symbols(naming, [spl144_276])])).
fof(f2888, plain, (spl144_275 <=> sP79), introduced(avatar_definition, [new_symbols(naming, [spl144_275])])).
fof(f2880, plain, (spl144_274 <=> sP80), introduced(avatar_definition, [new_symbols(naming, [spl144_274])])).
fof(f2872, plain, (spl144_273 <=> sP81), introduced(avatar_definition, [new_symbols(naming, [spl144_273])])).
fof(f2864, plain, (spl144_272 <=> sP82), introduced(avatar_definition, [new_symbols(naming, [spl144_272])])).
fof(f2856, plain, (spl144_271 <=> sP83), introduced(avatar_definition, [new_symbols(naming, [spl144_271])])).
fof(f2848, plain, (spl144_270 <=> sP84), introduced(avatar_definition, [new_symbols(naming, [spl144_270])])).
fof(f2840, plain, (spl144_269 <=> sP85), introduced(avatar_definition, [new_symbols(naming, [spl144_269])])).
fof(f2832, plain, (spl144_268 <=> sP86), introduced(avatar_definition, [new_symbols(naming, [spl144_268])])).
fof(f2824, plain, (spl144_267 <=> sP87), introduced(avatar_definition, [new_symbols(naming, [spl144_267])])).
fof(f2816, plain, (spl144_266 <=> sP88), introduced(avatar_definition, [new_symbols(naming, [spl144_266])])).
fof(f2808, plain, (spl144_265 <=> sP89), introduced(avatar_definition, [new_symbols(naming, [spl144_265])])).
fof(f2800, plain, (spl144_264 <=> sP90), introduced(avatar_definition, [new_symbols(naming, [spl144_264])])).
fof(f2792, plain, (spl144_263 <=> sP91), introduced(avatar_definition, [new_symbols(naming, [spl144_263])])).
fof(f2784, plain, (spl144_262 <=> sP92), introduced(avatar_definition, [new_symbols(naming, [spl144_262])])).
fof(f2776, plain, (spl144_261 <=> sP93), introduced(avatar_definition, [new_symbols(naming, [spl144_261])])).
fof(f2768, plain, (spl144_260 <=> sP94), introduced(avatar_definition, [new_symbols(naming, [spl144_260])])).
fof(f2760, plain, (spl144_259 <=> sP95), introduced(avatar_definition, [new_symbols(naming, [spl144_259])])).
fof(f2752, plain, (spl144_258 <=> sP96), introduced(avatar_definition, [new_symbols(naming, [spl144_258])])).
fof(f2744, plain, (spl144_257 <=> sP97), introduced(avatar_definition, [new_symbols(naming, [spl144_257])])).
fof(f2736, plain, (spl144_256 <=> sP98), introduced(avatar_definition, [new_symbols(naming, [spl144_256])])).
fof(f2728, plain, (spl144_255 <=> sP99), introduced(avatar_definition, [new_symbols(naming, [spl144_255])])).
fof(f2720, plain, (spl144_254 <=> sP100), introduced(avatar_definition, [new_symbols(naming, [spl144_254])])).
fof(f2712, plain, (spl144_253 <=> sP101), introduced(avatar_definition, [new_symbols(naming, [spl144_253])])).
fof(f2704, plain, (spl144_252 <=> sP102), introduced(avatar_definition, [new_symbols(naming, [spl144_252])])).
fof(f2696, plain, (spl144_251 <=> sP103), introduced(avatar_definition, [new_symbols(naming, [spl144_251])])).
fof(f2688, plain, (spl144_250 <=> sP104), introduced(avatar_definition, [new_symbols(naming, [spl144_250])])).
fof(f2680, plain, (spl144_249 <=> sP105), introduced(avatar_definition, [new_symbols(naming, [spl144_249])])).
fof(f2672, plain, (spl144_248 <=> sP106), introduced(avatar_definition, [new_symbols(naming, [spl144_248])])).
fof(f2664, plain, (spl144_247 <=> sP107), introduced(avatar_definition, [new_symbols(naming, [spl144_247])])).
fof(f2656, plain, (spl144_246 <=> sP108), introduced(avatar_definition, [new_symbols(naming, [spl144_246])])).
fof(f2648, plain, (spl144_245 <=> sP109), introduced(avatar_definition, [new_symbols(naming, [spl144_245])])).
fof(f2640, plain, (spl144_244 <=> sP110), introduced(avatar_definition, [new_symbols(naming, [spl144_244])])).
fof(f2632, plain, (spl144_243 <=> sP111), introduced(avatar_definition, [new_symbols(naming, [spl144_243])])).
fof(f2624, plain, (spl144_242 <=> sP112), introduced(avatar_definition, [new_symbols(naming, [spl144_242])])).
fof(f2616, plain, (spl144_241 <=> sP113), introduced(avatar_definition, [new_symbols(naming, [spl144_241])])).
fof(f2608, plain, (spl144_240 <=> sP114), introduced(avatar_definition, [new_symbols(naming, [spl144_240])])).
fof(f2600, plain, (spl144_239 <=> sP115), introduced(avatar_definition, [new_symbols(naming, [spl144_239])])).
fof(f2592, plain, (spl144_238 <=> sP116), introduced(avatar_definition, [new_symbols(naming, [spl144_238])])).
fof(f2584, plain, (spl144_237 <=> sP117), introduced(avatar_definition, [new_symbols(naming, [spl144_237])])).
fof(f2576, plain, (spl144_236 <=> sP118), introduced(avatar_definition, [new_symbols(naming, [spl144_236])])).
fof(f2568, plain, (spl144_235 <=> sP119), introduced(avatar_definition, [new_symbols(naming, [spl144_235])])).
fof(f2560, plain, (spl144_234 <=> sP120), introduced(avatar_definition, [new_symbols(naming, [spl144_234])])).
fof(f2552, plain, (spl144_233 <=> sP121), introduced(avatar_definition, [new_symbols(naming, [spl144_233])])).
fof(f2544, plain, (spl144_232 <=> sP122), introduced(avatar_definition, [new_symbols(naming, [spl144_232])])).
fof(f2536, plain, (spl144_231 <=> sP123), introduced(avatar_definition, [new_symbols(naming, [spl144_231])])).
fof(f2528, plain, (spl144_230 <=> sP124), introduced(avatar_definition, [new_symbols(naming, [spl144_230])])).
fof(f2520, plain, (spl144_229 <=> sP125), introduced(avatar_definition, [new_symbols(naming, [spl144_229])])).
fof(f2512, plain, (spl144_228 <=> sP126), introduced(avatar_definition, [new_symbols(naming, [spl144_228])])).
fof(f2504, plain, (spl144_227 <=> sP127), introduced(avatar_definition, [new_symbols(naming, [spl144_227])])).
fof(f2496, plain, (spl144_226 <=> sP128), introduced(avatar_definition, [new_symbols(naming, [spl144_226])])).
fof(f1096, plain, ((e23 = op2(e23, e23)) | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66), inference(cnf_transformation, [], [f154])).
fof(f154, plain, (((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e23, op2(e23, e23))) & (e22 = op2(e23, op2(e23, e22))) & (e21 = op2(e23, op2(e23, e21))) & (e20 = op2(e23, op2(e23, e20)))) | sP131 | sP130 | sP129) & ((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66)), inference(definition_folding, [], [f11, e153, e152, e151, e150, e149, e148, e147, e146, e145, e144, e143, e142, e141, e140, e139, e138, e137, e136, e135, e134, e133, e132, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88])).
fof(f88, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP66), inference(usedef, [], [e88])).
fof(e88, plain, (sP66 <=> (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f89, plain, ((~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP67), inference(usedef, [], [e89])).
fof(e89, plain, (sP67 <=> (~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f90, plain, ((~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP68), inference(usedef, [], [e90])).
fof(e90, plain, (sP68 <=> (~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f91, plain, ((~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP69), inference(usedef, [], [e91])).
fof(e91, plain, (sP69 <=> (~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f92, plain, ((~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | ~ sP70), inference(usedef, [], [e92])).
fof(e92, plain, (sP70 <=> (~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f93, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | ~ sP71), inference(usedef, [], [e93])).
fof(e93, plain, (sP71 <=> (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f94, plain, ((~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | ~ sP72), inference(usedef, [], [e94])).
fof(e94, plain, (sP72 <=> (~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f95, plain, ((~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | ~ sP73), inference(usedef, [], [e95])).
fof(e95, plain, (sP73 <=> (~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f96, plain, ((~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | ~ sP74), inference(usedef, [], [e96])).
fof(e96, plain, (sP74 <=> (~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f97, plain, ((~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | ~ sP75), inference(usedef, [], [e97])).
fof(e97, plain, (sP75 <=> (~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f98, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | ~ sP76), inference(usedef, [], [e98])).
fof(e98, plain, (sP76 <=> (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f99, plain, ((~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | ~ sP77), inference(usedef, [], [e99])).
fof(e99, plain, (sP77 <=> (~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f100, plain, ((~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | ~ sP78), inference(usedef, [], [e100])).
fof(e100, plain, (sP78 <=> (~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f101, plain, ((~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | ~ sP79), inference(usedef, [], [e101])).
fof(e101, plain, (sP79 <=> (~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f102, plain, ((~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | ~ sP80), inference(usedef, [], [e102])).
fof(e102, plain, (sP80 <=> (~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f103, plain, ((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | ~ sP81), inference(usedef, [], [e103])).
fof(e103, plain, (sP81 <=> (~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f104, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | ~ sP82), inference(usedef, [], [e104])).
fof(e104, plain, (sP82 <=> (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f105, plain, ((~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | ~ sP83), inference(usedef, [], [e105])).
fof(e105, plain, (sP83 <=> (~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f106, plain, ((~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | ~ sP84), inference(usedef, [], [e106])).
fof(e106, plain, (sP84 <=> (~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f107, plain, ((~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | ~ sP85), inference(usedef, [], [e107])).
fof(e107, plain, (sP85 <=> (~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f108, plain, ((~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP86), inference(usedef, [], [e108])).
fof(e108, plain, (sP86 <=> (~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f109, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP87), inference(usedef, [], [e109])).
fof(e109, plain, (sP87 <=> (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f110, plain, ((~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP88), inference(usedef, [], [e110])).
fof(e110, plain, (sP88 <=> (~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f111, plain, ((~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP89), inference(usedef, [], [e111])).
fof(e111, plain, (sP89 <=> (~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f112, plain, ((~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | ~ sP90), inference(usedef, [], [e112])).
fof(e112, plain, (sP90 <=> (~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f113, plain, ((~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | ~ sP91), inference(usedef, [], [e113])).
fof(e113, plain, (sP91 <=> (~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f114, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | ~ sP92), inference(usedef, [], [e114])).
fof(e114, plain, (sP92 <=> (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f115, plain, ((~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | ~ sP93), inference(usedef, [], [e115])).
fof(e115, plain, (sP93 <=> (~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f116, plain, ((~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | ~ sP94), inference(usedef, [], [e116])).
fof(e116, plain, (sP94 <=> (~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f117, plain, ((~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | ~ sP95), inference(usedef, [], [e117])).
fof(e117, plain, (sP95 <=> (~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f118, plain, ((~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | ~ sP96), inference(usedef, [], [e118])).
fof(e118, plain, (sP96 <=> (~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f119, plain, ((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | ~ sP97), inference(usedef, [], [e119])).
fof(e119, plain, (sP97 <=> (~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f120, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | ~ sP98), inference(usedef, [], [e120])).
fof(e120, plain, (sP98 <=> (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f121, plain, ((~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | ~ sP99), inference(usedef, [], [e121])).
fof(e121, plain, (sP99 <=> (~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f122, plain, ((~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | ~ sP100), inference(usedef, [], [e122])).
fof(e122, plain, (sP100 <=> (~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f123, plain, ((~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | ~ sP101), inference(usedef, [], [e123])).
fof(e123, plain, (sP101 <=> (~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f124, plain, ((~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | ~ sP102), inference(usedef, [], [e124])).
fof(e124, plain, (sP102 <=> (~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f125, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | ~ sP103), inference(usedef, [], [e125])).
fof(e125, plain, (sP103 <=> (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f126, plain, ((~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | ~ sP104), inference(usedef, [], [e126])).
fof(e126, plain, (sP104 <=> (~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f127, plain, ((~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | ~ sP105), inference(usedef, [], [e127])).
fof(e127, plain, (sP105 <=> (~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f128, plain, ((~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP106), inference(usedef, [], [e128])).
fof(e128, plain, (sP106 <=> (~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f129, plain, ((~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP107), inference(usedef, [], [e129])).
fof(e129, plain, (sP107 <=> (~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f130, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP108), inference(usedef, [], [e130])).
fof(e130, plain, (sP108 <=> (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f131, plain, ((~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP109), inference(usedef, [], [e131])).
fof(e131, plain, (sP109 <=> (~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f132, plain, ((~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | ~ sP110), inference(usedef, [], [e132])).
fof(e132, plain, (sP110 <=> (~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f133, plain, ((~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | ~ sP111), inference(usedef, [], [e133])).
fof(e133, plain, (sP111 <=> (~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f134, plain, ((~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | ~ sP112), inference(usedef, [], [e134])).
fof(e134, plain, (sP112 <=> (~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f135, plain, ((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | ~ sP113), inference(usedef, [], [e135])).
fof(e135, plain, (sP113 <=> (~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f136, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | ~ sP114), inference(usedef, [], [e136])).
fof(e136, plain, (sP114 <=> (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f137, plain, ((~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | ~ sP115), inference(usedef, [], [e137])).
fof(e137, plain, (sP115 <=> (~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f138, plain, ((~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | ~ sP116), inference(usedef, [], [e138])).
fof(e138, plain, (sP116 <=> (~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f139, plain, ((~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | ~ sP117), inference(usedef, [], [e139])).
fof(e139, plain, (sP117 <=> (~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f140, plain, ((~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | ~ sP118), inference(usedef, [], [e140])).
fof(e140, plain, (sP118 <=> (~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f141, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | ~ sP119), inference(usedef, [], [e141])).
fof(e141, plain, (sP119 <=> (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f142, plain, ((~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | ~ sP120), inference(usedef, [], [e142])).
fof(e142, plain, (sP120 <=> (~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f143, plain, ((~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | ~ sP121), inference(usedef, [], [e143])).
fof(e143, plain, (sP121 <=> (~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f144, plain, ((~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | ~ sP122), inference(usedef, [], [e144])).
fof(e144, plain, (sP122 <=> (~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f145, plain, ((~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | ~ sP123), inference(usedef, [], [e145])).
fof(e145, plain, (sP123 <=> (~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f146, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | ~ sP124), inference(usedef, [], [e146])).
fof(e146, plain, (sP124 <=> (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f147, plain, ((~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | ~ sP125), inference(usedef, [], [e147])).
fof(e147, plain, (sP125 <=> (~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f148, plain, ((~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP126), inference(usedef, [], [e148])).
fof(e148, plain, (sP126 <=> (~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f149, plain, ((~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP127), inference(usedef, [], [e149])).
fof(e149, plain, (sP127 <=> (~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f150, plain, ((~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP128), inference(usedef, [], [e150])).
fof(e150, plain, (sP128 <=> (~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f151, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, op2(e20, e23))) & (e22 = op2(e20, op2(e20, e22))) & (e21 = op2(e20, op2(e20, e21))) & (e20 = op2(e20, op2(e20, e20)))) | ~ sP129), inference(usedef, [], [e151])).
fof(e151, plain, (sP129 <=> (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, op2(e20, e23))) & (e22 = op2(e20, op2(e20, e22))) & (e21 = op2(e20, op2(e20, e21))) & (e20 = op2(e20, op2(e20, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f152, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e23 = op2(e21, op2(e21, e23))) & (e22 = op2(e21, op2(e21, e22))) & (e21 = op2(e21, op2(e21, e21))) & (e20 = op2(e21, op2(e21, e20)))) | ~ sP130), inference(usedef, [], [e152])).
fof(e152, plain, (sP130 <=> (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e23 = op2(e21, op2(e21, e23))) & (e22 = op2(e21, op2(e21, e22))) & (e21 = op2(e21, op2(e21, e21))) & (e20 = op2(e21, op2(e21, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP130])])).
fof(f153, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e23 = op2(e22, op2(e22, e23))) & (e22 = op2(e22, op2(e22, e22))) & (e21 = op2(e22, op2(e22, e21))) & (e20 = op2(e22, op2(e22, e20)))) | ~ sP131), inference(usedef, [], [e153])).
fof(e153, plain, (sP131 <=> (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e23 = op2(e22, op2(e22, e23))) & (e22 = op2(e22, op2(e22, e22))) & (e21 = op2(e22, op2(e22, e21))) & (e20 = op2(e22, op2(e22, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP131])])).
fof(f11, plain, (((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e23, op2(e23, e23))) & (e22 = op2(e23, op2(e23, e22))) & (e21 = op2(e23, op2(e23, e21))) & (e20 = op2(e23, op2(e23, e20)))) | (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e23 = op2(e22, op2(e22, e23))) & (e22 = op2(e22, op2(e22, e22))) & (e21 = op2(e22, op2(e22, e21))) & (e20 = op2(e22, op2(e22, e20)))) | (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e23 = op2(e21, op2(e21, e23))) & (e22 = op2(e21, op2(e21, e22))) & (e21 = op2(e21, op2(e21, e21))) & (e20 = op2(e21, op2(e21, e20)))) | (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, op2(e20, e23))) & (e22 = op2(e20, op2(e20, e22))) & (e21 = op2(e20, op2(e20, e21))) & (e20 = op2(e20, op2(e20, e20))))) & ((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | (~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | (~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | (~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | (~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | (~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | (~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | (~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | (~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | (~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | (~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | (~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | (~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | (~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | (~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | (~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | (~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | (~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | (~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | (~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | (~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | (~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | (~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | (~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | (~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | (~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | (~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | (~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | (~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | (~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | (~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | (~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | (~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | (~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | (~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | (~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | (~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | (~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | (~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | (~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | (~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | (~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | (~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | (~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | (~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | (~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | (~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | (~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | (~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | (~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | (~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | (~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | (~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | (~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | (~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax11)).
fof(f3021, plain, (spl144_221 | spl144_216 | spl144_211 | spl144_292), inference(avatar_split_clause, [], [f1100, f3018, f2415, f2442, f2469])).
fof(f2469, plain, (spl144_221 <=> sP129), introduced(avatar_definition, [new_symbols(naming, [spl144_221])])).
fof(f2442, plain, (spl144_216 <=> sP130), introduced(avatar_definition, [new_symbols(naming, [spl144_216])])).
fof(f2415, plain, (spl144_211 <=> sP131), introduced(avatar_definition, [new_symbols(naming, [spl144_211])])).
fof(f1100, plain, ((e20 = op2(e23, op2(e23, e20))) | sP131 | sP130 | sP129), inference(cnf_transformation, [], [f154])).
fof(f3011, plain, (spl144_221 | spl144_216 | spl144_211 | spl144_290), inference(avatar_split_clause, [], [f1102, f3008, f2415, f2442, f2469])).
fof(f1102, plain, ((e22 = op2(e23, op2(e23, e22))) | sP131 | sP130 | sP129), inference(cnf_transformation, [], [f154])).
fof(f3000, plain, (spl144_221 | spl144_216 | spl144_211 | ~ spl144_108 | spl144_98), inference(avatar_split_clause, [], [f1105, f1639, f1681, f2415, f2442, f2469])).
fof(f1105, plain, ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21)) | sP131 | sP130 | sP129), inference(cnf_transformation, [], [f154])).
fof(f2998, plain, (~ spl144_288 | spl144_125), inference(avatar_split_clause, [], [f1092, f1754, f2992])).
fof(f1092, plain, ((e20 = op2(e20, e20)) | ~ sP66), inference(cnf_transformation, [], [f299])).
fof(f299, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP66), inference(nnf_transformation, [], [f88])).
fof(f2990, plain, (~ spl144_287 | spl144_125), inference(avatar_split_clause, [], [f1088, f1754, f2984])).
fof(f1088, plain, ((e20 = op2(e20, e20)) | ~ sP67), inference(cnf_transformation, [], [f298])).
fof(f298, plain, ((~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP67), inference(nnf_transformation, [], [f89])).
fof(f2982, plain, (~ spl144_286 | spl144_125), inference(avatar_split_clause, [], [f1084, f1754, f2976])).
fof(f1084, plain, ((e20 = op2(e20, e20)) | ~ sP68), inference(cnf_transformation, [], [f297])).
fof(f297, plain, ((~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP68), inference(nnf_transformation, [], [f90])).
fof(f2974, plain, (~ spl144_285 | spl144_125), inference(avatar_split_clause, [], [f1080, f1754, f2968])).
fof(f1080, plain, ((e20 = op2(e20, e20)) | ~ sP69), inference(cnf_transformation, [], [f296])).
fof(f296, plain, ((~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP69), inference(nnf_transformation, [], [f91])).
fof(f2964, plain, (~ spl144_284 | spl144_126), inference(avatar_split_clause, [], [f1078, f1758, f2960])).
fof(f1078, plain, ((op2(e20, e20) = e21) | ~ sP70), inference(cnf_transformation, [], [f295])).
fof(f295, plain, ((~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | ~ sP70), inference(nnf_transformation, [], [f92])).
fof(f2956, plain, (~ spl144_283 | spl144_106), inference(avatar_split_clause, [], [f1074, f1673, f2952])).
fof(f1074, plain, ((e21 = op2(e21, e21)) | ~ sP71), inference(cnf_transformation, [], [f294])).
fof(f294, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | ~ sP71), inference(nnf_transformation, [], [f93])).
fof(f2955, plain, (~ spl144_283 | ~ spl144_106), inference(avatar_split_clause, [], [f1075, f1673, f2952])).
fof(f1075, plain, (~ (e21 = op2(e21, e21)) | ~ sP71), inference(cnf_transformation, [], [f294])).
fof(f2950, plain, (~ spl144_282 | spl144_121), inference(avatar_split_clause, [], [f1068, f1737, f2944])).
fof(f1068, plain, ((e20 = op2(e20, e21)) | ~ sP72), inference(cnf_transformation, [], [f293])).
fof(f293, plain, ((~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | ~ sP72), inference(nnf_transformation, [], [f94])).
fof(f2948, plain, (~ spl144_282 | spl144_86), inference(avatar_split_clause, [], [f1070, f1588, f2944])).
fof(f1070, plain, ((e21 = op2(e22, e22)) | ~ sP72), inference(cnf_transformation, [], [f293])).
fof(f2947, plain, (~ spl144_282 | ~ spl144_91), inference(avatar_split_clause, [], [f1071, f1609, f2944])).
fof(f1071, plain, (~ (e22 = op2(e22, e21)) | ~ sP72), inference(cnf_transformation, [], [f293])).
fof(f2942, plain, (~ spl144_281 | spl144_121), inference(avatar_split_clause, [], [f1064, f1737, f2936])).
fof(f1064, plain, ((e20 = op2(e20, e21)) | ~ sP73), inference(cnf_transformation, [], [f292])).
fof(f292, plain, ((~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (op2(e20, e20) = e21) & (e20 = op2(e20, e21))) | ~ sP73), inference(nnf_transformation, [], [f95])).
fof(f2939, plain, (~ spl144_281 | ~ spl144_76), inference(avatar_split_clause, [], [f1067, f1545, f2936])).
fof(f1067, plain, (~ (e23 = op2(e23, e21)) | ~ sP73), inference(cnf_transformation, [], [f292])).
fof(f2932, plain, (~ spl144_280 | spl144_127), inference(avatar_split_clause, [], [f1062, f1762, f2928])).
fof(f1062, plain, ((op2(e20, e20) = e22) | ~ sP74), inference(cnf_transformation, [], [f291])).
fof(f291, plain, ((~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | ~ sP74), inference(nnf_transformation, [], [f96])).
fof(f2926, plain, (~ spl144_279 | spl144_117), inference(avatar_split_clause, [], [f1056, f1720, f2920])).
fof(f1056, plain, ((e20 = op2(e20, e22)) | ~ sP75), inference(cnf_transformation, [], [f290])).
fof(f290, plain, ((~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | ~ sP75), inference(nnf_transformation, [], [f97])).
fof(f2924, plain, (~ spl144_279 | spl144_107), inference(avatar_split_clause, [], [f1058, f1677, f2920])).
fof(f1058, plain, ((e22 = op2(e21, e21)) | ~ sP75), inference(cnf_transformation, [], [f290])).
fof(f2923, plain, (~ spl144_279 | ~ spl144_102), inference(avatar_split_clause, [], [f1059, f1656, f2920])).
fof(f1059, plain, (~ (e21 = op2(e21, e22)) | ~ sP75), inference(cnf_transformation, [], [f290])).
fof(f2918, plain, (~ spl144_278 | spl144_117), inference(avatar_split_clause, [], [f1052, f1720, f2912])).
fof(f1052, plain, ((e20 = op2(e20, e22)) | ~ sP76), inference(cnf_transformation, [], [f289])).
fof(f289, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | ~ sP76), inference(nnf_transformation, [], [f98])).
fof(f2916, plain, (~ spl144_278 | spl144_87), inference(avatar_split_clause, [], [f1054, f1592, f2912])).
fof(f1054, plain, ((e22 = op2(e22, e22)) | ~ sP76), inference(cnf_transformation, [], [f289])).
fof(f2915, plain, (~ spl144_278 | ~ spl144_87), inference(avatar_split_clause, [], [f1055, f1592, f2912])).
fof(f1055, plain, (~ (e22 = op2(e22, e22)) | ~ sP76), inference(cnf_transformation, [], [f289])).
fof(f2908, plain, (~ spl144_277 | spl144_67), inference(avatar_split_clause, [], [f1050, f1507, f2904])).
fof(f1050, plain, ((e22 = op2(e23, e23)) | ~ sP77), inference(cnf_transformation, [], [f288])).
fof(f288, plain, ((~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (op2(e20, e20) = e22) & (e20 = op2(e20, e22))) | ~ sP77), inference(nnf_transformation, [], [f99])).
fof(f2901, plain, (~ spl144_276 | ~ spl144_128), inference(avatar_split_clause, [], [f1045, f1766, f2896])).
fof(f1045, plain, (~ (op2(e20, e20) = e23) | ~ sP78), inference(cnf_transformation, [], [f287])).
fof(f287, plain, ((~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | ~ sP78), inference(nnf_transformation, [], [f100])).
fof(f2893, plain, (~ spl144_275 | ~ spl144_128), inference(avatar_split_clause, [], [f1041, f1766, f2888])).
fof(f1041, plain, (~ (op2(e20, e20) = e23) | ~ sP79), inference(cnf_transformation, [], [f286])).
fof(f286, plain, ((~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | ~ sP79), inference(nnf_transformation, [], [f101])).
fof(f2885, plain, (~ spl144_274 | ~ spl144_128), inference(avatar_split_clause, [], [f1037, f1766, f2880])).
fof(f1037, plain, (~ (op2(e20, e20) = e23) | ~ sP80), inference(cnf_transformation, [], [f285])).
fof(f285, plain, ((~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | ~ sP80), inference(nnf_transformation, [], [f102])).
fof(f2877, plain, (~ spl144_273 | ~ spl144_128), inference(avatar_split_clause, [], [f1033, f1766, f2872])).
fof(f1033, plain, (~ (op2(e20, e20) = e23) | ~ sP81), inference(cnf_transformation, [], [f284])).
fof(f284, plain, ((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (op2(e20, e20) = e23) & (e20 = op2(e20, e23))) | ~ sP81), inference(nnf_transformation, [], [f103])).
fof(f2868, plain, (~ spl144_272 | spl144_125), inference(avatar_split_clause, [], [f1030, f1754, f2864])).
fof(f1030, plain, ((e20 = op2(e20, e20)) | ~ sP82), inference(cnf_transformation, [], [f283])).
fof(f283, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | ~ sP82), inference(nnf_transformation, [], [f104])).
fof(f2861, plain, (~ spl144_271 | ~ spl144_105), inference(avatar_split_clause, [], [f1025, f1669, f2856])).
fof(f1025, plain, (~ (e20 = op2(e21, e21)) | ~ sP83), inference(cnf_transformation, [], [f282])).
fof(f282, plain, ((~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | ~ sP83), inference(nnf_transformation, [], [f105])).
fof(f2860, plain, (~ spl144_271 | spl144_105), inference(avatar_split_clause, [], [f1026, f1669, f2856])).
fof(f1026, plain, ((e20 = op2(e21, e21)) | ~ sP83), inference(cnf_transformation, [], [f282])).
fof(f2854, plain, (~ spl144_270 | spl144_110), inference(avatar_split_clause, [], [f1020, f1690, f2848])).
fof(f1020, plain, ((e21 = op2(e21, e20)) | ~ sP84), inference(cnf_transformation, [], [f281])).
fof(f281, plain, ((~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | ~ sP84), inference(nnf_transformation, [], [f106])).
fof(f2852, plain, (~ spl144_270 | spl144_85), inference(avatar_split_clause, [], [f1022, f1584, f2848])).
fof(f1022, plain, ((e20 = op2(e22, e22)) | ~ sP84), inference(cnf_transformation, [], [f281])).
fof(f2851, plain, (~ spl144_270 | ~ spl144_95), inference(avatar_split_clause, [], [f1023, f1626, f2848])).
fof(f1023, plain, (~ (e22 = op2(e22, e20)) | ~ sP84), inference(cnf_transformation, [], [f281])).
fof(f2844, plain, (~ spl144_269 | spl144_65), inference(avatar_split_clause, [], [f1018, f1499, f2840])).
fof(f1018, plain, ((e20 = op2(e23, e23)) | ~ sP85), inference(cnf_transformation, [], [f280])).
fof(f280, plain, ((~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e21, e21)) & (e21 = op2(e21, e20))) | ~ sP85), inference(nnf_transformation, [], [f107])).
fof(f2836, plain, (~ spl144_268 | spl144_126), inference(avatar_split_clause, [], [f1014, f1758, f2832])).
fof(f1014, plain, ((op2(e20, e20) = e21) | ~ sP86), inference(cnf_transformation, [], [f279])).
fof(f279, plain, ((~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP86), inference(nnf_transformation, [], [f108])).
fof(f2830, plain, (~ spl144_267 | spl144_106), inference(avatar_split_clause, [], [f1008, f1673, f2824])).
fof(f1008, plain, ((e21 = op2(e21, e21)) | ~ sP87), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP87), inference(nnf_transformation, [], [f109])).
fof(f2829, plain, (~ spl144_267 | ~ spl144_106), inference(avatar_split_clause, [], [f1009, f1673, f2824])).
fof(f1009, plain, (~ (e21 = op2(e21, e21)) | ~ sP87), inference(cnf_transformation, [], [f278])).
fof(f2822, plain, (~ spl144_266 | spl144_106), inference(avatar_split_clause, [], [f1004, f1673, f2816])).
fof(f1004, plain, ((e21 = op2(e21, e21)) | ~ sP88), inference(cnf_transformation, [], [f277])).
fof(f277, plain, ((~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP88), inference(nnf_transformation, [], [f110])).
fof(f2821, plain, (~ spl144_266 | ~ spl144_106), inference(avatar_split_clause, [], [f1005, f1673, f2816])).
fof(f1005, plain, (~ (e21 = op2(e21, e21)) | ~ sP88), inference(cnf_transformation, [], [f277])).
fof(f2814, plain, (~ spl144_265 | spl144_106), inference(avatar_split_clause, [], [f1000, f1673, f2808])).
fof(f1000, plain, ((e21 = op2(e21, e21)) | ~ sP89), inference(cnf_transformation, [], [f276])).
fof(f276, plain, ((~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP89), inference(nnf_transformation, [], [f111])).
fof(f2813, plain, (~ spl144_265 | ~ spl144_106), inference(avatar_split_clause, [], [f1001, f1673, f2808])).
fof(f1001, plain, (~ (e21 = op2(e21, e21)) | ~ sP89), inference(cnf_transformation, [], [f276])).
fof(f2804, plain, (~ spl144_264 | spl144_127), inference(avatar_split_clause, [], [f998, f1762, f2800])).
fof(f998, plain, ((op2(e20, e20) = e22) | ~ sP90), inference(cnf_transformation, [], [f275])).
fof(f275, plain, ((~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | ~ sP90), inference(nnf_transformation, [], [f112])).
fof(f2797, plain, (~ spl144_263 | ~ spl144_107), inference(avatar_split_clause, [], [f993, f1677, f2792])).
fof(f993, plain, (~ (e22 = op2(e21, e21)) | ~ sP91), inference(cnf_transformation, [], [f274])).
fof(f274, plain, ((~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | ~ sP91), inference(nnf_transformation, [], [f113])).
fof(f2796, plain, (~ spl144_263 | spl144_107), inference(avatar_split_clause, [], [f994, f1677, f2792])).
fof(f994, plain, ((e22 = op2(e21, e21)) | ~ sP91), inference(cnf_transformation, [], [f274])).
fof(f2788, plain, (~ spl144_262 | spl144_87), inference(avatar_split_clause, [], [f990, f1592, f2784])).
fof(f990, plain, ((e22 = op2(e22, e22)) | ~ sP92), inference(cnf_transformation, [], [f273])).
fof(f273, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | ~ sP92), inference(nnf_transformation, [], [f114])).
fof(f2787, plain, (~ spl144_262 | ~ spl144_87), inference(avatar_split_clause, [], [f991, f1592, f2784])).
fof(f991, plain, (~ (e22 = op2(e22, e22)) | ~ sP92), inference(cnf_transformation, [], [f273])).
fof(f2780, plain, (~ spl144_261 | spl144_67), inference(avatar_split_clause, [], [f986, f1507, f2776])).
fof(f986, plain, ((e22 = op2(e23, e23)) | ~ sP93), inference(cnf_transformation, [], [f272])).
fof(f272, plain, ((~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e21, e21)) & (e21 = op2(e21, e22))) | ~ sP93), inference(nnf_transformation, [], [f115])).
fof(f2774, plain, (~ spl144_260 | spl144_98), inference(avatar_split_clause, [], [f980, f1639, f2768])).
fof(f980, plain, ((e21 = op2(e21, e23)) | ~ sP94), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ((~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | ~ sP94), inference(nnf_transformation, [], [f116])).
fof(f2766, plain, (~ spl144_259 | spl144_98), inference(avatar_split_clause, [], [f976, f1639, f2760])).
fof(f976, plain, ((e21 = op2(e21, e23)) | ~ sP95), inference(cnf_transformation, [], [f270])).
fof(f270, plain, ((~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | ~ sP95), inference(nnf_transformation, [], [f117])).
fof(f2758, plain, (~ spl144_258 | spl144_98), inference(avatar_split_clause, [], [f972, f1639, f2752])).
fof(f972, plain, ((e21 = op2(e21, e23)) | ~ sP96), inference(cnf_transformation, [], [f269])).
fof(f269, plain, ((~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | ~ sP96), inference(nnf_transformation, [], [f118])).
fof(f2750, plain, (~ spl144_257 | spl144_98), inference(avatar_split_clause, [], [f968, f1639, f2744])).
fof(f968, plain, ((e21 = op2(e21, e23)) | ~ sP97), inference(cnf_transformation, [], [f268])).
fof(f268, plain, ((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (e23 = op2(e21, e21)) & (e21 = op2(e21, e23))) | ~ sP97), inference(nnf_transformation, [], [f119])).
fof(f2740, plain, (~ spl144_256 | spl144_125), inference(avatar_split_clause, [], [f966, f1754, f2736])).
fof(f966, plain, ((e20 = op2(e20, e20)) | ~ sP98), inference(cnf_transformation, [], [f267])).
fof(f267, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | ~ sP98), inference(nnf_transformation, [], [f120])).
fof(f2734, plain, (~ spl144_255 | spl144_95), inference(avatar_split_clause, [], [f960, f1626, f2728])).
fof(f960, plain, ((e22 = op2(e22, e20)) | ~ sP99), inference(cnf_transformation, [], [f266])).
fof(f266, plain, ((~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | ~ sP99), inference(nnf_transformation, [], [f121])).
fof(f2732, plain, (~ spl144_255 | spl144_105), inference(avatar_split_clause, [], [f962, f1669, f2728])).
fof(f962, plain, ((e20 = op2(e21, e21)) | ~ sP99), inference(cnf_transformation, [], [f266])).
fof(f2731, plain, (~ spl144_255 | ~ spl144_110), inference(avatar_split_clause, [], [f963, f1690, f2728])).
fof(f963, plain, (~ (e21 = op2(e21, e20)) | ~ sP99), inference(cnf_transformation, [], [f266])).
fof(f2726, plain, (~ spl144_254 | spl144_95), inference(avatar_split_clause, [], [f956, f1626, f2720])).
fof(f956, plain, ((e22 = op2(e22, e20)) | ~ sP100), inference(cnf_transformation, [], [f265])).
fof(f265, plain, ((~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | ~ sP100), inference(nnf_transformation, [], [f122])).
fof(f2725, plain, (~ spl144_254 | ~ spl144_85), inference(avatar_split_clause, [], [f957, f1584, f2720])).
fof(f957, plain, (~ (e20 = op2(e22, e22)) | ~ sP100), inference(cnf_transformation, [], [f265])).
fof(f2724, plain, (~ spl144_254 | spl144_85), inference(avatar_split_clause, [], [f958, f1584, f2720])).
fof(f958, plain, ((e20 = op2(e22, e22)) | ~ sP100), inference(cnf_transformation, [], [f265])).
fof(f2716, plain, (~ spl144_253 | spl144_65), inference(avatar_split_clause, [], [f954, f1499, f2712])).
fof(f954, plain, ((e20 = op2(e23, e23)) | ~ sP101), inference(cnf_transformation, [], [f264])).
fof(f264, plain, ((~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e22, e22)) & (e22 = op2(e22, e20))) | ~ sP101), inference(nnf_transformation, [], [f123])).
fof(f2708, plain, (~ spl144_252 | spl144_126), inference(avatar_split_clause, [], [f950, f1758, f2704])).
fof(f950, plain, ((op2(e20, e20) = e21) | ~ sP102), inference(cnf_transformation, [], [f263])).
fof(f263, plain, ((~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | ~ sP102), inference(nnf_transformation, [], [f124])).
fof(f2700, plain, (~ spl144_251 | spl144_106), inference(avatar_split_clause, [], [f946, f1673, f2696])).
fof(f946, plain, ((e21 = op2(e21, e21)) | ~ sP103), inference(cnf_transformation, [], [f262])).
fof(f262, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | ~ sP103), inference(nnf_transformation, [], [f125])).
fof(f2699, plain, (~ spl144_251 | ~ spl144_106), inference(avatar_split_clause, [], [f947, f1673, f2696])).
fof(f947, plain, (~ (e21 = op2(e21, e21)) | ~ sP103), inference(cnf_transformation, [], [f262])).
fof(f2693, plain, (~ spl144_250 | ~ spl144_86), inference(avatar_split_clause, [], [f941, f1588, f2688])).
fof(f941, plain, (~ (e21 = op2(e22, e22)) | ~ sP104), inference(cnf_transformation, [], [f261])).
fof(f261, plain, ((~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | ~ sP104), inference(nnf_transformation, [], [f126])).
fof(f2692, plain, (~ spl144_250 | spl144_86), inference(avatar_split_clause, [], [f942, f1588, f2688])).
fof(f942, plain, ((e21 = op2(e22, e22)) | ~ sP104), inference(cnf_transformation, [], [f261])).
fof(f2686, plain, (~ spl144_249 | spl144_91), inference(avatar_split_clause, [], [f936, f1609, f2680])).
fof(f936, plain, ((e22 = op2(e22, e21)) | ~ sP105), inference(cnf_transformation, [], [f260])).
fof(f260, plain, ((~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e22, e22)) & (e22 = op2(e22, e21))) | ~ sP105), inference(nnf_transformation, [], [f127])).
fof(f2685, plain, (~ spl144_249 | ~ spl144_86), inference(avatar_split_clause, [], [f937, f1588, f2680])).
fof(f937, plain, (~ (e21 = op2(e22, e22)) | ~ sP105), inference(cnf_transformation, [], [f260])).
fof(f2683, plain, (~ spl144_249 | ~ spl144_76), inference(avatar_split_clause, [], [f939, f1545, f2680])).
fof(f939, plain, (~ (e23 = op2(e23, e21)) | ~ sP105), inference(cnf_transformation, [], [f260])).
fof(f2676, plain, (~ spl144_248 | spl144_127), inference(avatar_split_clause, [], [f934, f1762, f2672])).
fof(f934, plain, ((op2(e20, e20) = e22) | ~ sP106), inference(cnf_transformation, [], [f259])).
fof(f259, plain, ((~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP106), inference(nnf_transformation, [], [f128])).
fof(f2670, plain, (~ spl144_247 | spl144_87), inference(avatar_split_clause, [], [f928, f1592, f2664])).
fof(f928, plain, ((e22 = op2(e22, e22)) | ~ sP107), inference(cnf_transformation, [], [f258])).
fof(f258, plain, ((~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP107), inference(nnf_transformation, [], [f129])).
fof(f2669, plain, (~ spl144_247 | ~ spl144_87), inference(avatar_split_clause, [], [f929, f1592, f2664])).
fof(f929, plain, (~ (e22 = op2(e22, e22)) | ~ sP107), inference(cnf_transformation, [], [f258])).
fof(f2662, plain, (~ spl144_246 | spl144_87), inference(avatar_split_clause, [], [f924, f1592, f2656])).
fof(f924, plain, ((e22 = op2(e22, e22)) | ~ sP108), inference(cnf_transformation, [], [f257])).
fof(f257, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP108), inference(nnf_transformation, [], [f130])).
fof(f2661, plain, (~ spl144_246 | ~ spl144_87), inference(avatar_split_clause, [], [f925, f1592, f2656])).
fof(f925, plain, (~ (e22 = op2(e22, e22)) | ~ sP108), inference(cnf_transformation, [], [f257])).
fof(f2652, plain, (~ spl144_245 | spl144_67), inference(avatar_split_clause, [], [f922, f1507, f2648])).
fof(f922, plain, ((e22 = op2(e23, e23)) | ~ sP109), inference(cnf_transformation, [], [f256])).
fof(f256, plain, ((~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP109), inference(nnf_transformation, [], [f131])).
fof(f2646, plain, (~ spl144_244 | spl144_83), inference(avatar_split_clause, [], [f916, f1575, f2640])).
fof(f916, plain, ((e22 = op2(e22, e23)) | ~ sP110), inference(cnf_transformation, [], [f255])).
fof(f255, plain, ((~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | ~ sP110), inference(nnf_transformation, [], [f132])).
fof(f2645, plain, (~ spl144_244 | ~ spl144_88), inference(avatar_split_clause, [], [f917, f1596, f2640])).
fof(f917, plain, (~ (e23 = op2(e22, e22)) | ~ sP110), inference(cnf_transformation, [], [f255])).
fof(f2643, plain, (~ spl144_244 | ~ spl144_113), inference(avatar_split_clause, [], [f919, f1703, f2640])).
fof(f919, plain, (~ (e20 = op2(e20, e23)) | ~ sP110), inference(cnf_transformation, [], [f255])).
fof(f2638, plain, (~ spl144_243 | spl144_83), inference(avatar_split_clause, [], [f912, f1575, f2632])).
fof(f912, plain, ((e22 = op2(e22, e23)) | ~ sP111), inference(cnf_transformation, [], [f254])).
fof(f254, plain, ((~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | ~ sP111), inference(nnf_transformation, [], [f133])).
fof(f2637, plain, (~ spl144_243 | ~ spl144_88), inference(avatar_split_clause, [], [f913, f1596, f2632])).
fof(f913, plain, (~ (e23 = op2(e22, e22)) | ~ sP111), inference(cnf_transformation, [], [f254])).
fof(f2636, plain, (~ spl144_243 | spl144_108), inference(avatar_split_clause, [], [f914, f1681, f2632])).
fof(f914, plain, ((e23 = op2(e21, e21)) | ~ sP111), inference(cnf_transformation, [], [f254])).
fof(f2630, plain, (~ spl144_242 | spl144_83), inference(avatar_split_clause, [], [f908, f1575, f2624])).
fof(f908, plain, ((e22 = op2(e22, e23)) | ~ sP112), inference(cnf_transformation, [], [f253])).
fof(f253, plain, ((~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | ~ sP112), inference(nnf_transformation, [], [f134])).
fof(f2629, plain, (~ spl144_242 | ~ spl144_88), inference(avatar_split_clause, [], [f909, f1596, f2624])).
fof(f909, plain, (~ (e23 = op2(e22, e22)) | ~ sP112), inference(cnf_transformation, [], [f253])).
fof(f2628, plain, (~ spl144_242 | spl144_88), inference(avatar_split_clause, [], [f910, f1596, f2624])).
fof(f910, plain, ((e23 = op2(e22, e22)) | ~ sP112), inference(cnf_transformation, [], [f253])).
fof(f2620, plain, (~ spl144_241 | spl144_68), inference(avatar_split_clause, [], [f906, f1511, f2616])).
fof(f906, plain, ((e23 = op2(e23, e23)) | ~ sP113), inference(cnf_transformation, [], [f252])).
fof(f252, plain, ((~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23)) & ~ (e23 = op2(e22, e22)) & (e22 = op2(e22, e23))) | ~ sP113), inference(nnf_transformation, [], [f135])).
fof(f2614, plain, (~ spl144_240 | spl144_80), inference(avatar_split_clause, [], [f900, f1562, f2608])).
fof(f900, plain, ((e23 = op2(e23, e20)) | ~ sP114), inference(cnf_transformation, [], [f251])).
fof(f251, plain, ((~ (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | ~ sP114), inference(nnf_transformation, [], [f136])).
fof(f2606, plain, (~ spl144_239 | spl144_80), inference(avatar_split_clause, [], [f896, f1562, f2600])).
fof(f896, plain, ((e23 = op2(e23, e20)) | ~ sP115), inference(cnf_transformation, [], [f250])).
fof(f250, plain, ((~ (e21 = op2(e21, e20)) & (e20 = op2(e21, e21)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | ~ sP115), inference(nnf_transformation, [], [f137])).
fof(f2598, plain, (~ spl144_238 | spl144_80), inference(avatar_split_clause, [], [f892, f1562, f2592])).
fof(f892, plain, ((e23 = op2(e23, e20)) | ~ sP116), inference(cnf_transformation, [], [f249])).
fof(f249, plain, ((~ (e22 = op2(e22, e20)) & (e20 = op2(e22, e22)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | ~ sP116), inference(nnf_transformation, [], [f138])).
fof(f2590, plain, (~ spl144_237 | spl144_80), inference(avatar_split_clause, [], [f888, f1562, f2584])).
fof(f888, plain, ((e23 = op2(e23, e20)) | ~ sP117), inference(cnf_transformation, [], [f248])).
fof(f248, plain, ((~ (e23 = op2(e23, e20)) & (e20 = op2(e23, e23)) & ~ (e20 = op2(e23, e23)) & (e23 = op2(e23, e20))) | ~ sP117), inference(nnf_transformation, [], [f139])).
fof(f2581, plain, (~ spl144_236 | ~ spl144_66), inference(avatar_split_clause, [], [f885, f1503, f2576])).
fof(f885, plain, (~ (e21 = op2(e23, e23)) | ~ sP118), inference(cnf_transformation, [], [f247])).
fof(f247, plain, ((~ (e20 = op2(e20, e21)) & (op2(e20, e20) = e21) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | ~ sP118), inference(nnf_transformation, [], [f140])).
fof(f2573, plain, (~ spl144_235 | ~ spl144_66), inference(avatar_split_clause, [], [f881, f1503, f2568])).
fof(f881, plain, (~ (e21 = op2(e23, e23)) | ~ sP119), inference(cnf_transformation, [], [f246])).
fof(f246, plain, ((~ (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | ~ sP119), inference(nnf_transformation, [], [f141])).
fof(f2565, plain, (~ spl144_234 | ~ spl144_66), inference(avatar_split_clause, [], [f877, f1503, f2560])).
fof(f877, plain, (~ (e21 = op2(e23, e23)) | ~ sP120), inference(cnf_transformation, [], [f245])).
fof(f245, plain, ((~ (e22 = op2(e22, e21)) & (e21 = op2(e22, e22)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | ~ sP120), inference(nnf_transformation, [], [f142])).
fof(f2557, plain, (~ spl144_233 | ~ spl144_66), inference(avatar_split_clause, [], [f873, f1503, f2552])).
fof(f873, plain, (~ (e21 = op2(e23, e23)) | ~ sP121), inference(cnf_transformation, [], [f244])).
fof(f244, plain, ((~ (e23 = op2(e23, e21)) & (e21 = op2(e23, e23)) & ~ (e21 = op2(e23, e23)) & (e23 = op2(e23, e21))) | ~ sP121), inference(nnf_transformation, [], [f143])).
fof(f2548, plain, (~ spl144_232 | spl144_127), inference(avatar_split_clause, [], [f870, f1762, f2544])).
fof(f870, plain, ((op2(e20, e20) = e22) | ~ sP122), inference(cnf_transformation, [], [f243])).
fof(f243, plain, ((~ (e20 = op2(e20, e22)) & (op2(e20, e20) = e22) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | ~ sP122), inference(nnf_transformation, [], [f144])).
fof(f2542, plain, (~ spl144_231 | spl144_72), inference(avatar_split_clause, [], [f864, f1528, f2536])).
fof(f864, plain, ((e23 = op2(e23, e22)) | ~ sP123), inference(cnf_transformation, [], [f242])).
fof(f242, plain, ((~ (e21 = op2(e21, e22)) & (e22 = op2(e21, e21)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | ~ sP123), inference(nnf_transformation, [], [f145])).
fof(f2540, plain, (~ spl144_231 | spl144_107), inference(avatar_split_clause, [], [f866, f1677, f2536])).
fof(f866, plain, ((e22 = op2(e21, e21)) | ~ sP123), inference(cnf_transformation, [], [f242])).
fof(f2539, plain, (~ spl144_231 | ~ spl144_102), inference(avatar_split_clause, [], [f867, f1656, f2536])).
fof(f867, plain, (~ (e21 = op2(e21, e22)) | ~ sP123), inference(cnf_transformation, [], [f242])).
fof(f2534, plain, (~ spl144_230 | spl144_72), inference(avatar_split_clause, [], [f860, f1528, f2528])).
fof(f860, plain, ((e23 = op2(e23, e22)) | ~ sP124), inference(cnf_transformation, [], [f241])).
fof(f241, plain, ((~ (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | ~ sP124), inference(nnf_transformation, [], [f146])).
fof(f2532, plain, (~ spl144_230 | spl144_87), inference(avatar_split_clause, [], [f862, f1592, f2528])).
fof(f862, plain, ((e22 = op2(e22, e22)) | ~ sP124), inference(cnf_transformation, [], [f241])).
fof(f2531, plain, (~ spl144_230 | ~ spl144_87), inference(avatar_split_clause, [], [f863, f1592, f2528])).
fof(f863, plain, (~ (e22 = op2(e22, e22)) | ~ sP124), inference(cnf_transformation, [], [f241])).
fof(f2524, plain, (~ spl144_229 | spl144_67), inference(avatar_split_clause, [], [f858, f1507, f2520])).
fof(f858, plain, ((e22 = op2(e23, e23)) | ~ sP125), inference(cnf_transformation, [], [f240])).
fof(f240, plain, ((~ (e23 = op2(e23, e22)) & (e22 = op2(e23, e23)) & ~ (e22 = op2(e23, e23)) & (e23 = op2(e23, e22))) | ~ sP125), inference(nnf_transformation, [], [f147])).
fof(f2518, plain, (~ spl144_228 | spl144_68), inference(avatar_split_clause, [], [f852, f1511, f2512])).
fof(f852, plain, ((e23 = op2(e23, e23)) | ~ sP126), inference(cnf_transformation, [], [f239])).
fof(f239, plain, ((~ (e20 = op2(e20, e23)) & (op2(e20, e20) = e23) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP126), inference(nnf_transformation, [], [f148])).
fof(f2510, plain, (~ spl144_227 | spl144_68), inference(avatar_split_clause, [], [f848, f1511, f2504])).
fof(f848, plain, ((e23 = op2(e23, e23)) | ~ sP127), inference(cnf_transformation, [], [f238])).
fof(f238, plain, ((~ (e21 = op2(e21, e23)) & (e23 = op2(e21, e21)) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP127), inference(nnf_transformation, [], [f149])).
fof(f2502, plain, (~ spl144_226 | spl144_68), inference(avatar_split_clause, [], [f844, f1511, f2496])).
fof(f844, plain, ((e23 = op2(e23, e23)) | ~ sP128), inference(cnf_transformation, [], [f237])).
fof(f237, plain, ((~ (e22 = op2(e22, e23)) & (e23 = op2(e22, e22)) & ~ (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ~ sP128), inference(nnf_transformation, [], [f150])).
fof(f2494, plain, (~ spl144_221 | spl144_225), inference(avatar_split_clause, [], [f836, f2491, f2469])).
fof(f836, plain, ((e20 = op2(e20, op2(e20, e20))) | ~ sP129), inference(cnf_transformation, [], [f236])).
fof(f236, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, op2(e20, e23))) & (e22 = op2(e20, op2(e20, e22))) & (e21 = op2(e20, op2(e20, e21))) & (e20 = op2(e20, op2(e20, e20)))) | ~ sP129), inference(nnf_transformation, [], [f151])).
fof(f2489, plain, (~ spl144_221 | spl144_224), inference(avatar_split_clause, [], [f837, f2486, f2469])).
fof(f837, plain, ((e21 = op2(e20, op2(e20, e21))) | ~ sP129), inference(cnf_transformation, [], [f236])).
fof(f2479, plain, (~ spl144_221 | spl144_222), inference(avatar_split_clause, [], [f839, f2476, f2469])).
fof(f839, plain, ((e23 = op2(e20, op2(e20, e23))) | ~ sP129), inference(cnf_transformation, [], [f236])).
fof(f2473, plain, (~ spl144_221 | ~ spl144_85 | spl144_95), inference(avatar_split_clause, [], [f842, f1626, f1584, f2469])).
fof(f842, plain, ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22)) | ~ sP129), inference(cnf_transformation, [], [f236])).
fof(f2445, plain, (~ spl144_216 | ~ spl144_66 | spl144_76), inference(avatar_split_clause, [], [f835, f1545, f1503, f2442])).
fof(f835, plain, ((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23)) | ~ sP130), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e23 = op2(e21, op2(e21, e23))) & (e22 = op2(e21, op2(e21, e22))) & (e21 = op2(e21, op2(e21, e21))) & (e20 = op2(e21, op2(e21, e20)))) | ~ sP130), inference(nnf_transformation, [], [f152])).
fof(f2440, plain, (~ spl144_211 | spl144_215), inference(avatar_split_clause, [], [f820, f2437, f2415])).
fof(f820, plain, ((e20 = op2(e22, op2(e22, e20))) | ~ sP131), inference(cnf_transformation, [], [f234])).
fof(f234, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e23 = op2(e22, op2(e22, e23))) & (e22 = op2(e22, op2(e22, e22))) & (e21 = op2(e22, op2(e22, e21))) & (e20 = op2(e22, op2(e22, e20)))) | ~ sP131), inference(nnf_transformation, [], [f153])).
fof(f2435, plain, (~ spl144_211 | spl144_214), inference(avatar_split_clause, [], [f821, f2432, f2415])).
fof(f821, plain, ((e21 = op2(e22, op2(e22, e21))) | ~ sP131), inference(cnf_transformation, [], [f234])).
fof(f2425, plain, (~ spl144_211 | spl144_212), inference(avatar_split_clause, [], [f823, f2422, f2415])).
fof(f823, plain, ((e23 = op2(e22, op2(e22, e23))) | ~ sP131), inference(cnf_transformation, [], [f234])).
fof(f2419, plain, (~ spl144_211 | ~ spl144_107 | spl144_102), inference(avatar_split_clause, [], [f825, f1656, f1677, f2415])).
fof(f825, plain, ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21)) | ~ sP131), inference(cnf_transformation, [], [f234])).
fof(f2413, plain, (spl144_206 | spl144_205 | spl144_204 | spl144_203 | spl144_202 | spl144_201 | spl144_200 | spl144_199 | spl144_198 | spl144_197 | spl144_196 | spl144_195 | spl144_194 | spl144_193 | spl144_192 | spl144_191 | spl144_190 | spl144_189 | spl144_188 | spl144_187 | spl144_186 | spl144_185 | spl144_184 | spl144_183 | spl144_182 | spl144_181 | spl144_180 | spl144_179 | spl144_178 | spl144_177 | spl144_176 | spl144_175 | spl144_174 | spl144_173 | spl144_172 | spl144_171 | spl144_170 | spl144_169 | spl144_168 | spl144_167 | spl144_166 | spl144_165 | spl144_164 | spl144_163 | spl144_162 | spl144_161 | spl144_160 | spl144_159 | spl144_158 | spl144_157 | spl144_156 | spl144_155 | spl144_154 | spl144_153 | spl144_152 | spl144_151 | spl144_150 | spl144_149 | spl144_148 | spl144_147 | spl144_146 | spl144_145 | spl144_144 | spl144_4), inference(avatar_split_clause, [], [f808, f1207, f1884, f1892, f1900, f1908, f1916, f1924, f1932, f1940, f1948, f1956, f1964, f1972, f1980, f1988, f1996, f2004, f2012, f2020, f2028, f2036, f2044, f2052, f2060, f2068, f2076, f2084, f2092, f2100, f2108, f2116, f2124, f2132, f2140, f2148, f2156, f2164, f2172, f2180, f2188, f2196, f2204, f2212, f2220, f2228, f2236, f2244, f2252, f2260, f2268, f2276, f2284, f2292, f2300, f2308, f2316, f2324, f2332, f2340, f2348, f2356, f2364, f2372, f2380])).
fof(f2380, plain, (spl144_206 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl144_206])])).
fof(f2372, plain, (spl144_205 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl144_205])])).
fof(f2364, plain, (spl144_204 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl144_204])])).
fof(f2356, plain, (spl144_203 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl144_203])])).
fof(f2348, plain, (spl144_202 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl144_202])])).
fof(f2340, plain, (spl144_201 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl144_201])])).
fof(f2332, plain, (spl144_200 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl144_200])])).
fof(f2324, plain, (spl144_199 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl144_199])])).
fof(f2316, plain, (spl144_198 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl144_198])])).
fof(f2308, plain, (spl144_197 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl144_197])])).
fof(f2300, plain, (spl144_196 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl144_196])])).
fof(f2292, plain, (spl144_195 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl144_195])])).
fof(f2284, plain, (spl144_194 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl144_194])])).
fof(f2276, plain, (spl144_193 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl144_193])])).
fof(f2268, plain, (spl144_192 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl144_192])])).
fof(f2260, plain, (spl144_191 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl144_191])])).
fof(f2252, plain, (spl144_190 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl144_190])])).
fof(f2244, plain, (spl144_189 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl144_189])])).
fof(f2236, plain, (spl144_188 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl144_188])])).
fof(f2228, plain, (spl144_187 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl144_187])])).
fof(f2220, plain, (spl144_186 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl144_186])])).
fof(f2212, plain, (spl144_185 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl144_185])])).
fof(f2204, plain, (spl144_184 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl144_184])])).
fof(f2196, plain, (spl144_183 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl144_183])])).
fof(f2188, plain, (spl144_182 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl144_182])])).
fof(f2180, plain, (spl144_181 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl144_181])])).
fof(f2172, plain, (spl144_180 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl144_180])])).
fof(f2164, plain, (spl144_179 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl144_179])])).
fof(f2156, plain, (spl144_178 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl144_178])])).
fof(f2148, plain, (spl144_177 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl144_177])])).
fof(f2140, plain, (spl144_176 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl144_176])])).
fof(f2132, plain, (spl144_175 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl144_175])])).
fof(f2124, plain, (spl144_174 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl144_174])])).
fof(f2116, plain, (spl144_173 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl144_173])])).
fof(f2108, plain, (spl144_172 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl144_172])])).
fof(f2100, plain, (spl144_171 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl144_171])])).
fof(f2092, plain, (spl144_170 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl144_170])])).
fof(f2084, plain, (spl144_169 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl144_169])])).
fof(f2076, plain, (spl144_168 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl144_168])])).
fof(f2068, plain, (spl144_167 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl144_167])])).
fof(f2060, plain, (spl144_166 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl144_166])])).
fof(f2052, plain, (spl144_165 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl144_165])])).
fof(f2044, plain, (spl144_164 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl144_164])])).
fof(f2036, plain, (spl144_163 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl144_163])])).
fof(f2028, plain, (spl144_162 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl144_162])])).
fof(f2020, plain, (spl144_161 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl144_161])])).
fof(f2012, plain, (spl144_160 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl144_160])])).
fof(f2004, plain, (spl144_159 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl144_159])])).
fof(f1996, plain, (spl144_158 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl144_158])])).
fof(f1988, plain, (spl144_157 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl144_157])])).
fof(f1980, plain, (spl144_156 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl144_156])])).
fof(f1972, plain, (spl144_155 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl144_155])])).
fof(f1964, plain, (spl144_154 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl144_154])])).
fof(f1956, plain, (spl144_153 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl144_153])])).
fof(f1948, plain, (spl144_152 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl144_152])])).
fof(f1940, plain, (spl144_151 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl144_151])])).
fof(f1932, plain, (spl144_150 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl144_150])])).
fof(f1924, plain, (spl144_149 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl144_149])])).
fof(f1916, plain, (spl144_148 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl144_148])])).
fof(f1908, plain, (spl144_147 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl144_147])])).
fof(f1900, plain, (spl144_146 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl144_146])])).
fof(f1892, plain, (spl144_145 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl144_145])])).
fof(f1884, plain, (spl144_144 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl144_144])])).
fof(f808, plain, ((e13 = op1(e13, e13)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f87])).
fof(f87, plain, (((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e13, op1(e13, e13))) & (e12 = op1(e13, op1(e13, e12))) & (e11 = op1(e13, op1(e13, e11))) & (e10 = op1(e13, op1(e13, e10)))) | sP65 | sP64 | sP63) & ((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f10, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21])).
fof(f21, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP0), inference(usedef, [], [e21])).
fof(e21, plain, (sP0 <=> (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f22, plain, ((~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP1), inference(usedef, [], [e22])).
fof(e22, plain, (sP1 <=> (~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f23, plain, ((~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP2), inference(usedef, [], [e23])).
fof(e23, plain, (sP2 <=> (~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f24, plain, ((~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP3), inference(usedef, [], [e24])).
fof(e24, plain, (sP3 <=> (~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f25, plain, ((~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | ~ sP4), inference(usedef, [], [e25])).
fof(e25, plain, (sP4 <=> (~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f26, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | ~ sP5), inference(usedef, [], [e26])).
fof(e26, plain, (sP5 <=> (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f27, plain, ((~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | ~ sP6), inference(usedef, [], [e27])).
fof(e27, plain, (sP6 <=> (~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f28, plain, ((~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | ~ sP7), inference(usedef, [], [e28])).
fof(e28, plain, (sP7 <=> (~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f29, plain, ((~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | ~ sP8), inference(usedef, [], [e29])).
fof(e29, plain, (sP8 <=> (~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f30, plain, ((~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | ~ sP9), inference(usedef, [], [e30])).
fof(e30, plain, (sP9 <=> (~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f31, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | ~ sP10), inference(usedef, [], [e31])).
fof(e31, plain, (sP10 <=> (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f32, plain, ((~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | ~ sP11), inference(usedef, [], [e32])).
fof(e32, plain, (sP11 <=> (~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f33, plain, ((~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | ~ sP12), inference(usedef, [], [e33])).
fof(e33, plain, (sP12 <=> (~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f34, plain, ((~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | ~ sP13), inference(usedef, [], [e34])).
fof(e34, plain, (sP13 <=> (~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f35, plain, ((~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | ~ sP14), inference(usedef, [], [e35])).
fof(e35, plain, (sP14 <=> (~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f36, plain, ((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | ~ sP15), inference(usedef, [], [e36])).
fof(e36, plain, (sP15 <=> (~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f37, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | ~ sP16), inference(usedef, [], [e37])).
fof(e37, plain, (sP16 <=> (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f38, plain, ((~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | ~ sP17), inference(usedef, [], [e38])).
fof(e38, plain, (sP17 <=> (~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f39, plain, ((~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | ~ sP18), inference(usedef, [], [e39])).
fof(e39, plain, (sP18 <=> (~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f40, plain, ((~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | ~ sP19), inference(usedef, [], [e40])).
fof(e40, plain, (sP19 <=> (~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f41, plain, ((~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP20), inference(usedef, [], [e41])).
fof(e41, plain, (sP20 <=> (~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f42, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP21), inference(usedef, [], [e42])).
fof(e42, plain, (sP21 <=> (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f43, plain, ((~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP22), inference(usedef, [], [e43])).
fof(e43, plain, (sP22 <=> (~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f44, plain, ((~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP23), inference(usedef, [], [e44])).
fof(e44, plain, (sP23 <=> (~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f45, plain, ((~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | ~ sP24), inference(usedef, [], [e45])).
fof(e45, plain, (sP24 <=> (~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f46, plain, ((~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | ~ sP25), inference(usedef, [], [e46])).
fof(e46, plain, (sP25 <=> (~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f47, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | ~ sP26), inference(usedef, [], [e47])).
fof(e47, plain, (sP26 <=> (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f48, plain, ((~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | ~ sP27), inference(usedef, [], [e48])).
fof(e48, plain, (sP27 <=> (~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f49, plain, ((~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | ~ sP28), inference(usedef, [], [e49])).
fof(e49, plain, (sP28 <=> (~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f50, plain, ((~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | ~ sP29), inference(usedef, [], [e50])).
fof(e50, plain, (sP29 <=> (~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f51, plain, ((~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | ~ sP30), inference(usedef, [], [e51])).
fof(e51, plain, (sP30 <=> (~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f52, plain, ((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | ~ sP31), inference(usedef, [], [e52])).
fof(e52, plain, (sP31 <=> (~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f53, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | ~ sP32), inference(usedef, [], [e53])).
fof(e53, plain, (sP32 <=> (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f54, plain, ((~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | ~ sP33), inference(usedef, [], [e54])).
fof(e54, plain, (sP33 <=> (~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f55, plain, ((~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | ~ sP34), inference(usedef, [], [e55])).
fof(e55, plain, (sP34 <=> (~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f56, plain, ((~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | ~ sP35), inference(usedef, [], [e56])).
fof(e56, plain, (sP35 <=> (~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f57, plain, ((~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | ~ sP36), inference(usedef, [], [e57])).
fof(e57, plain, (sP36 <=> (~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f58, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | ~ sP37), inference(usedef, [], [e58])).
fof(e58, plain, (sP37 <=> (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f59, plain, ((~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | ~ sP38), inference(usedef, [], [e59])).
fof(e59, plain, (sP38 <=> (~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f60, plain, ((~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | ~ sP39), inference(usedef, [], [e60])).
fof(e60, plain, (sP39 <=> (~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f61, plain, ((~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP40), inference(usedef, [], [e61])).
fof(e61, plain, (sP40 <=> (~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f62, plain, ((~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP41), inference(usedef, [], [e62])).
fof(e62, plain, (sP41 <=> (~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f63, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP42), inference(usedef, [], [e63])).
fof(e63, plain, (sP42 <=> (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f64, plain, ((~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP43), inference(usedef, [], [e64])).
fof(e64, plain, (sP43 <=> (~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f65, plain, ((~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | ~ sP44), inference(usedef, [], [e65])).
fof(e65, plain, (sP44 <=> (~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f66, plain, ((~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | ~ sP45), inference(usedef, [], [e66])).
fof(e66, plain, (sP45 <=> (~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f67, plain, ((~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | ~ sP46), inference(usedef, [], [e67])).
fof(e67, plain, (sP46 <=> (~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f68, plain, ((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | ~ sP47), inference(usedef, [], [e68])).
fof(e68, plain, (sP47 <=> (~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f69, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | ~ sP48), inference(usedef, [], [e69])).
fof(e69, plain, (sP48 <=> (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f70, plain, ((~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | ~ sP49), inference(usedef, [], [e70])).
fof(e70, plain, (sP49 <=> (~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f71, plain, ((~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | ~ sP50), inference(usedef, [], [e71])).
fof(e71, plain, (sP50 <=> (~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f72, plain, ((~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | ~ sP51), inference(usedef, [], [e72])).
fof(e72, plain, (sP51 <=> (~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f73, plain, ((~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | ~ sP52), inference(usedef, [], [e73])).
fof(e73, plain, (sP52 <=> (~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f74, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | ~ sP53), inference(usedef, [], [e74])).
fof(e74, plain, (sP53 <=> (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f75, plain, ((~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | ~ sP54), inference(usedef, [], [e75])).
fof(e75, plain, (sP54 <=> (~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f76, plain, ((~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | ~ sP55), inference(usedef, [], [e76])).
fof(e76, plain, (sP55 <=> (~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f77, plain, ((~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | ~ sP56), inference(usedef, [], [e77])).
fof(e77, plain, (sP56 <=> (~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f78, plain, ((~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | ~ sP57), inference(usedef, [], [e78])).
fof(e78, plain, (sP57 <=> (~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f79, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | ~ sP58), inference(usedef, [], [e79])).
fof(e79, plain, (sP58 <=> (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f80, plain, ((~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | ~ sP59), inference(usedef, [], [e80])).
fof(e80, plain, (sP59 <=> (~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f81, plain, ((~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP60), inference(usedef, [], [e81])).
fof(e81, plain, (sP60 <=> (~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f82, plain, ((~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP61), inference(usedef, [], [e82])).
fof(e82, plain, (sP61 <=> (~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f83, plain, ((~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP62), inference(usedef, [], [e83])).
fof(e83, plain, (sP62 <=> (~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f84, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, op1(e10, e13))) & (e12 = op1(e10, op1(e10, e12))) & (e11 = op1(e10, op1(e10, e11))) & (e10 = op1(e10, op1(e10, e10)))) | ~ sP63), inference(usedef, [], [e84])).
fof(e84, plain, (sP63 <=> (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, op1(e10, e13))) & (e12 = op1(e10, op1(e10, e12))) & (e11 = op1(e10, op1(e10, e11))) & (e10 = op1(e10, op1(e10, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f85, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e13 = op1(e11, op1(e11, e13))) & (e12 = op1(e11, op1(e11, e12))) & (e11 = op1(e11, op1(e11, e11))) & (e10 = op1(e11, op1(e11, e10)))) | ~ sP64), inference(usedef, [], [e85])).
fof(e85, plain, (sP64 <=> (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e13 = op1(e11, op1(e11, e13))) & (e12 = op1(e11, op1(e11, e12))) & (e11 = op1(e11, op1(e11, e11))) & (e10 = op1(e11, op1(e11, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f86, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e13 = op1(e12, op1(e12, e13))) & (e12 = op1(e12, op1(e12, e12))) & (e11 = op1(e12, op1(e12, e11))) & (e10 = op1(e12, op1(e12, e10)))) | ~ sP65), inference(usedef, [], [e86])).
fof(e86, plain, (sP65 <=> (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e13 = op1(e12, op1(e12, e13))) & (e12 = op1(e12, op1(e12, e12))) & (e11 = op1(e12, op1(e12, e11))) & (e10 = op1(e12, op1(e12, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f10, plain, (((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e13, op1(e13, e13))) & (e12 = op1(e13, op1(e13, e12))) & (e11 = op1(e13, op1(e13, e11))) & (e10 = op1(e13, op1(e13, e10)))) | (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e13 = op1(e12, op1(e12, e13))) & (e12 = op1(e12, op1(e12, e12))) & (e11 = op1(e12, op1(e12, e11))) & (e10 = op1(e12, op1(e12, e10)))) | (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e13 = op1(e11, op1(e11, e13))) & (e12 = op1(e11, op1(e11, e12))) & (e11 = op1(e11, op1(e11, e11))) & (e10 = op1(e11, op1(e11, e10)))) | (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, op1(e10, e13))) & (e12 = op1(e10, op1(e10, e12))) & (e11 = op1(e10, op1(e10, e11))) & (e10 = op1(e10, op1(e10, e10))))) & ((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | (~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | (~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | (~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | (~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | (~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | (~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | (~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | (~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | (~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | (~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | (~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | (~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | (~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | (~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | (~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | (~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | (~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | (~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | (~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | (~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | (~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | (~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | (~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | (~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | (~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | (~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | (~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | (~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | (~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | (~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | (~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | (~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | (~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | (~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | (~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | (~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | (~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | (~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | (~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | (~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | (~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | (~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | (~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | (~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | (~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | (~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | (~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | (~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | (~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | (~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | (~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | (~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | (~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | (~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax10)).
fof(f2409, plain, (spl144_139 | spl144_134 | spl144_129 | spl144_210), inference(avatar_split_clause, [], [f812, f2406, f1803, f1830, f1857])).
fof(f1857, plain, (spl144_139 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl144_139])])).
fof(f1830, plain, (spl144_134 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl144_134])])).
fof(f1803, plain, (spl144_129 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl144_129])])).
fof(f812, plain, ((e10 = op1(e13, op1(e13, e10))) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f87])).
fof(f2386, plain, (~ spl144_206 | spl144_61), inference(avatar_split_clause, [], [f804, f1450, f2380])).
fof(f1450, plain, (spl144_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl144_61])])).
fof(f804, plain, ((e10 = op1(e10, e10)) | ~ sP0), inference(cnf_transformation, [], [f233])).
fof(f233, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP0), inference(nnf_transformation, [], [f21])).
fof(f2385, plain, (~ spl144_206 | ~ spl144_61), inference(avatar_split_clause, [], [f805, f1450, f2380])).
fof(f805, plain, (~ (e10 = op1(e10, e10)) | ~ sP0), inference(cnf_transformation, [], [f233])).
fof(f2378, plain, (~ spl144_205 | spl144_61), inference(avatar_split_clause, [], [f800, f1450, f2372])).
fof(f800, plain, ((e10 = op1(e10, e10)) | ~ sP1), inference(cnf_transformation, [], [f232])).
fof(f232, plain, ((~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP1), inference(nnf_transformation, [], [f22])).
fof(f2377, plain, (~ spl144_205 | ~ spl144_61), inference(avatar_split_clause, [], [f801, f1450, f2372])).
fof(f801, plain, (~ (e10 = op1(e10, e10)) | ~ sP1), inference(cnf_transformation, [], [f232])).
fof(f2370, plain, (~ spl144_204 | spl144_61), inference(avatar_split_clause, [], [f796, f1450, f2364])).
fof(f796, plain, ((e10 = op1(e10, e10)) | ~ sP2), inference(cnf_transformation, [], [f231])).
fof(f231, plain, ((~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP2), inference(nnf_transformation, [], [f23])).
fof(f2369, plain, (~ spl144_204 | ~ spl144_61), inference(avatar_split_clause, [], [f797, f1450, f2364])).
fof(f797, plain, (~ (e10 = op1(e10, e10)) | ~ sP2), inference(cnf_transformation, [], [f231])).
fof(f2360, plain, (~ spl144_203 | spl144_1), inference(avatar_split_clause, [], [f794, f1195, f2356])).
fof(f794, plain, ((e10 = op1(e13, e13)) | ~ sP3), inference(cnf_transformation, [], [f230])).
fof(f230, plain, ((~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP3), inference(nnf_transformation, [], [f24])).
fof(f2352, plain, (~ spl144_202 | spl144_62), inference(avatar_split_clause, [], [f790, f1454, f2348])).
fof(f790, plain, ((op1(e10, e10) = e11) | ~ sP4), inference(cnf_transformation, [], [f229])).
fof(f229, plain, ((~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | ~ sP4), inference(nnf_transformation, [], [f25])).
fof(f2344, plain, (~ spl144_201 | spl144_42), inference(avatar_split_clause, [], [f786, f1369, f2340])).
fof(f786, plain, ((e11 = op1(e11, e11)) | ~ sP5), inference(cnf_transformation, [], [f228])).
fof(f228, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | ~ sP5), inference(nnf_transformation, [], [f26])).
fof(f2343, plain, (~ spl144_201 | ~ spl144_42), inference(avatar_split_clause, [], [f787, f1369, f2340])).
fof(f787, plain, (~ (e11 = op1(e11, e11)) | ~ sP5), inference(cnf_transformation, [], [f228])).
fof(f2338, plain, (~ spl144_200 | spl144_57), inference(avatar_split_clause, [], [f780, f1433, f2332])).
fof(f780, plain, ((e10 = op1(e10, e11)) | ~ sP6), inference(cnf_transformation, [], [f227])).
fof(f227, plain, ((~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | ~ sP6), inference(nnf_transformation, [], [f27])).
fof(f2336, plain, (~ spl144_200 | spl144_22), inference(avatar_split_clause, [], [f782, f1284, f2332])).
fof(f782, plain, ((e11 = op1(e12, e12)) | ~ sP6), inference(cnf_transformation, [], [f227])).
fof(f2335, plain, (~ spl144_200 | ~ spl144_27), inference(avatar_split_clause, [], [f783, f1305, f2332])).
fof(f783, plain, (~ (e12 = op1(e12, e11)) | ~ sP6), inference(cnf_transformation, [], [f227])).
fof(f2330, plain, (~ spl144_199 | spl144_57), inference(avatar_split_clause, [], [f776, f1433, f2324])).
fof(f776, plain, ((e10 = op1(e10, e11)) | ~ sP7), inference(cnf_transformation, [], [f226])).
fof(f226, plain, ((~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (op1(e10, e10) = e11) & (e10 = op1(e10, e11))) | ~ sP7), inference(nnf_transformation, [], [f28])).
fof(f2327, plain, (~ spl144_199 | ~ spl144_12), inference(avatar_split_clause, [], [f779, f1241, f2324])).
fof(f779, plain, (~ (e13 = op1(e13, e11)) | ~ sP7), inference(cnf_transformation, [], [f226])).
fof(f2320, plain, (~ spl144_198 | spl144_63), inference(avatar_split_clause, [], [f774, f1458, f2316])).
fof(f774, plain, ((op1(e10, e10) = e12) | ~ sP8), inference(cnf_transformation, [], [f225])).
fof(f225, plain, ((~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | ~ sP8), inference(nnf_transformation, [], [f29])).
fof(f2314, plain, (~ spl144_197 | spl144_53), inference(avatar_split_clause, [], [f768, f1416, f2308])).
fof(f768, plain, ((e10 = op1(e10, e12)) | ~ sP9), inference(cnf_transformation, [], [f224])).
fof(f224, plain, ((~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | ~ sP9), inference(nnf_transformation, [], [f30])).
fof(f2312, plain, (~ spl144_197 | spl144_43), inference(avatar_split_clause, [], [f770, f1373, f2308])).
fof(f770, plain, ((e12 = op1(e11, e11)) | ~ sP9), inference(cnf_transformation, [], [f224])).
fof(f2311, plain, (~ spl144_197 | ~ spl144_38), inference(avatar_split_clause, [], [f771, f1352, f2308])).
fof(f771, plain, (~ (e11 = op1(e11, e12)) | ~ sP9), inference(cnf_transformation, [], [f224])).
fof(f2306, plain, (~ spl144_196 | spl144_53), inference(avatar_split_clause, [], [f764, f1416, f2300])).
fof(f764, plain, ((e10 = op1(e10, e12)) | ~ sP10), inference(cnf_transformation, [], [f223])).
fof(f223, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | ~ sP10), inference(nnf_transformation, [], [f31])).
fof(f2304, plain, (~ spl144_196 | spl144_23), inference(avatar_split_clause, [], [f766, f1288, f2300])).
fof(f766, plain, ((e12 = op1(e12, e12)) | ~ sP10), inference(cnf_transformation, [], [f223])).
fof(f2303, plain, (~ spl144_196 | ~ spl144_23), inference(avatar_split_clause, [], [f767, f1288, f2300])).
fof(f767, plain, (~ (e12 = op1(e12, e12)) | ~ sP10), inference(cnf_transformation, [], [f223])).
fof(f2298, plain, (~ spl144_195 | spl144_53), inference(avatar_split_clause, [], [f760, f1416, f2292])).
fof(f760, plain, ((e10 = op1(e10, e12)) | ~ sP11), inference(cnf_transformation, [], [f222])).
fof(f222, plain, ((~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (op1(e10, e10) = e12) & (e10 = op1(e10, e12))) | ~ sP11), inference(nnf_transformation, [], [f32])).
fof(f2296, plain, (~ spl144_195 | spl144_3), inference(avatar_split_clause, [], [f762, f1203, f2292])).
fof(f762, plain, ((e12 = op1(e13, e13)) | ~ sP11), inference(cnf_transformation, [], [f222])).
fof(f2295, plain, (~ spl144_195 | ~ spl144_8), inference(avatar_split_clause, [], [f763, f1224, f2292])).
fof(f763, plain, (~ (e13 = op1(e13, e12)) | ~ sP11), inference(cnf_transformation, [], [f222])).
fof(f2289, plain, (~ spl144_194 | ~ spl144_64), inference(avatar_split_clause, [], [f757, f1462, f2284])).
fof(f757, plain, (~ (op1(e10, e10) = e13) | ~ sP12), inference(cnf_transformation, [], [f221])).
fof(f221, plain, ((~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | ~ sP12), inference(nnf_transformation, [], [f33])).
fof(f2281, plain, (~ spl144_193 | ~ spl144_64), inference(avatar_split_clause, [], [f753, f1462, f2276])).
fof(f753, plain, (~ (op1(e10, e10) = e13) | ~ sP13), inference(cnf_transformation, [], [f220])).
fof(f220, plain, ((~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | ~ sP13), inference(nnf_transformation, [], [f34])).
fof(f2273, plain, (~ spl144_192 | ~ spl144_64), inference(avatar_split_clause, [], [f749, f1462, f2268])).
fof(f749, plain, (~ (op1(e10, e10) = e13) | ~ sP14), inference(cnf_transformation, [], [f219])).
fof(f219, plain, ((~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | ~ sP14), inference(nnf_transformation, [], [f35])).
fof(f2265, plain, (~ spl144_191 | ~ spl144_64), inference(avatar_split_clause, [], [f745, f1462, f2260])).
fof(f745, plain, (~ (op1(e10, e10) = e13) | ~ sP15), inference(cnf_transformation, [], [f218])).
fof(f218, plain, ((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (op1(e10, e10) = e13) & (e10 = op1(e10, e13))) | ~ sP15), inference(nnf_transformation, [], [f36])).
fof(f2258, plain, (~ spl144_190 | spl144_46), inference(avatar_split_clause, [], [f740, f1386, f2252])).
fof(f740, plain, ((e11 = op1(e11, e10)) | ~ sP16), inference(cnf_transformation, [], [f217])).
fof(f217, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | ~ sP16), inference(nnf_transformation, [], [f37])).
fof(f2256, plain, (~ spl144_190 | spl144_61), inference(avatar_split_clause, [], [f742, f1450, f2252])).
fof(f742, plain, ((e10 = op1(e10, e10)) | ~ sP16), inference(cnf_transformation, [], [f217])).
fof(f2255, plain, (~ spl144_190 | ~ spl144_61), inference(avatar_split_clause, [], [f743, f1450, f2252])).
fof(f743, plain, (~ (e10 = op1(e10, e10)) | ~ sP16), inference(cnf_transformation, [], [f217])).
fof(f2249, plain, (~ spl144_189 | ~ spl144_41), inference(avatar_split_clause, [], [f737, f1365, f2244])).
fof(f737, plain, (~ (e10 = op1(e11, e11)) | ~ sP17), inference(cnf_transformation, [], [f216])).
fof(f216, plain, ((~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | ~ sP17), inference(nnf_transformation, [], [f38])).
fof(f2248, plain, (~ spl144_189 | spl144_41), inference(avatar_split_clause, [], [f738, f1365, f2244])).
fof(f738, plain, ((e10 = op1(e11, e11)) | ~ sP17), inference(cnf_transformation, [], [f216])).
fof(f2242, plain, (~ spl144_188 | spl144_46), inference(avatar_split_clause, [], [f732, f1386, f2236])).
fof(f732, plain, ((e11 = op1(e11, e10)) | ~ sP18), inference(cnf_transformation, [], [f215])).
fof(f215, plain, ((~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | ~ sP18), inference(nnf_transformation, [], [f39])).
fof(f2240, plain, (~ spl144_188 | spl144_21), inference(avatar_split_clause, [], [f734, f1280, f2236])).
fof(f734, plain, ((e10 = op1(e12, e12)) | ~ sP18), inference(cnf_transformation, [], [f215])).
fof(f2232, plain, (~ spl144_187 | spl144_1), inference(avatar_split_clause, [], [f730, f1195, f2228])).
fof(f730, plain, ((e10 = op1(e13, e13)) | ~ sP19), inference(cnf_transformation, [], [f214])).
fof(f214, plain, ((~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e11, e11)) & (e11 = op1(e11, e10))) | ~ sP19), inference(nnf_transformation, [], [f40])).
fof(f2224, plain, (~ spl144_186 | spl144_62), inference(avatar_split_clause, [], [f726, f1454, f2220])).
fof(f726, plain, ((op1(e10, e10) = e11) | ~ sP20), inference(cnf_transformation, [], [f213])).
fof(f213, plain, ((~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP20), inference(nnf_transformation, [], [f41])).
fof(f2218, plain, (~ spl144_185 | spl144_42), inference(avatar_split_clause, [], [f720, f1369, f2212])).
fof(f720, plain, ((e11 = op1(e11, e11)) | ~ sP21), inference(cnf_transformation, [], [f212])).
fof(f212, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP21), inference(nnf_transformation, [], [f42])).
fof(f2217, plain, (~ spl144_185 | ~ spl144_42), inference(avatar_split_clause, [], [f721, f1369, f2212])).
fof(f721, plain, (~ (e11 = op1(e11, e11)) | ~ sP21), inference(cnf_transformation, [], [f212])).
fof(f2210, plain, (~ spl144_184 | spl144_42), inference(avatar_split_clause, [], [f716, f1369, f2204])).
fof(f716, plain, ((e11 = op1(e11, e11)) | ~ sP22), inference(cnf_transformation, [], [f211])).
fof(f211, plain, ((~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP22), inference(nnf_transformation, [], [f43])).
fof(f2209, plain, (~ spl144_184 | ~ spl144_42), inference(avatar_split_clause, [], [f717, f1369, f2204])).
fof(f717, plain, (~ (e11 = op1(e11, e11)) | ~ sP22), inference(cnf_transformation, [], [f211])).
fof(f2202, plain, (~ spl144_183 | spl144_42), inference(avatar_split_clause, [], [f712, f1369, f2196])).
fof(f712, plain, ((e11 = op1(e11, e11)) | ~ sP23), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ((~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP23), inference(nnf_transformation, [], [f44])).
fof(f2201, plain, (~ spl144_183 | ~ spl144_42), inference(avatar_split_clause, [], [f713, f1369, f2196])).
fof(f713, plain, (~ (e11 = op1(e11, e11)) | ~ sP23), inference(cnf_transformation, [], [f210])).
fof(f2192, plain, (~ spl144_182 | spl144_63), inference(avatar_split_clause, [], [f710, f1458, f2188])).
fof(f710, plain, ((op1(e10, e10) = e12) | ~ sP24), inference(cnf_transformation, [], [f209])).
fof(f209, plain, ((~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | ~ sP24), inference(nnf_transformation, [], [f45])).
fof(f2185, plain, (~ spl144_181 | ~ spl144_43), inference(avatar_split_clause, [], [f705, f1373, f2180])).
fof(f705, plain, (~ (e12 = op1(e11, e11)) | ~ sP25), inference(cnf_transformation, [], [f208])).
fof(f208, plain, ((~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | ~ sP25), inference(nnf_transformation, [], [f46])).
fof(f2184, plain, (~ spl144_181 | spl144_43), inference(avatar_split_clause, [], [f706, f1373, f2180])).
fof(f706, plain, ((e12 = op1(e11, e11)) | ~ sP25), inference(cnf_transformation, [], [f208])).
fof(f2178, plain, (~ spl144_180 | spl144_38), inference(avatar_split_clause, [], [f700, f1352, f2172])).
fof(f700, plain, ((e11 = op1(e11, e12)) | ~ sP26), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | ~ sP26), inference(nnf_transformation, [], [f47])).
fof(f2176, plain, (~ spl144_180 | spl144_23), inference(avatar_split_clause, [], [f702, f1288, f2172])).
fof(f702, plain, ((e12 = op1(e12, e12)) | ~ sP26), inference(cnf_transformation, [], [f207])).
fof(f2175, plain, (~ spl144_180 | ~ spl144_23), inference(avatar_split_clause, [], [f703, f1288, f2172])).
fof(f703, plain, (~ (e12 = op1(e12, e12)) | ~ sP26), inference(cnf_transformation, [], [f207])).
fof(f2170, plain, (~ spl144_179 | spl144_38), inference(avatar_split_clause, [], [f696, f1352, f2164])).
fof(f696, plain, ((e11 = op1(e11, e12)) | ~ sP27), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ((~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e11, e11)) & (e11 = op1(e11, e12))) | ~ sP27), inference(nnf_transformation, [], [f48])).
fof(f2169, plain, (~ spl144_179 | ~ spl144_43), inference(avatar_split_clause, [], [f697, f1373, f2164])).
fof(f697, plain, (~ (e12 = op1(e11, e11)) | ~ sP27), inference(cnf_transformation, [], [f206])).
fof(f2168, plain, (~ spl144_179 | spl144_3), inference(avatar_split_clause, [], [f698, f1203, f2164])).
fof(f698, plain, ((e12 = op1(e13, e13)) | ~ sP27), inference(cnf_transformation, [], [f206])).
fof(f2167, plain, (~ spl144_179 | ~ spl144_8), inference(avatar_split_clause, [], [f699, f1224, f2164])).
fof(f699, plain, (~ (e13 = op1(e13, e12)) | ~ sP27), inference(cnf_transformation, [], [f206])).
fof(f2162, plain, (~ spl144_178 | spl144_34), inference(avatar_split_clause, [], [f692, f1335, f2156])).
fof(f692, plain, ((e11 = op1(e11, e13)) | ~ sP28), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ((~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | ~ sP28), inference(nnf_transformation, [], [f49])).
fof(f2154, plain, (~ spl144_177 | spl144_34), inference(avatar_split_clause, [], [f688, f1335, f2148])).
fof(f688, plain, ((e11 = op1(e11, e13)) | ~ sP29), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ((~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | ~ sP29), inference(nnf_transformation, [], [f50])).
fof(f2146, plain, (~ spl144_176 | spl144_34), inference(avatar_split_clause, [], [f684, f1335, f2140])).
fof(f684, plain, ((e11 = op1(e11, e13)) | ~ sP30), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ((~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | ~ sP30), inference(nnf_transformation, [], [f51])).
fof(f2138, plain, (~ spl144_175 | spl144_34), inference(avatar_split_clause, [], [f680, f1335, f2132])).
fof(f680, plain, ((e11 = op1(e11, e13)) | ~ sP31), inference(cnf_transformation, [], [f202])).
fof(f202, plain, ((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (e13 = op1(e11, e11)) & (e11 = op1(e11, e13))) | ~ sP31), inference(nnf_transformation, [], [f52])).
fof(f2130, plain, (~ spl144_174 | spl144_31), inference(avatar_split_clause, [], [f676, f1322, f2124])).
fof(f676, plain, ((e12 = op1(e12, e10)) | ~ sP32), inference(cnf_transformation, [], [f201])).
fof(f201, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | ~ sP32), inference(nnf_transformation, [], [f53])).
fof(f2122, plain, (~ spl144_173 | spl144_31), inference(avatar_split_clause, [], [f672, f1322, f2116])).
fof(f672, plain, ((e12 = op1(e12, e10)) | ~ sP33), inference(cnf_transformation, [], [f200])).
fof(f200, plain, ((~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | ~ sP33), inference(nnf_transformation, [], [f54])).
fof(f2114, plain, (~ spl144_172 | spl144_31), inference(avatar_split_clause, [], [f668, f1322, f2108])).
fof(f668, plain, ((e12 = op1(e12, e10)) | ~ sP34), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ((~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | ~ sP34), inference(nnf_transformation, [], [f55])).
fof(f2106, plain, (~ spl144_171 | spl144_31), inference(avatar_split_clause, [], [f664, f1322, f2100])).
fof(f664, plain, ((e12 = op1(e12, e10)) | ~ sP35), inference(cnf_transformation, [], [f198])).
fof(f198, plain, ((~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e12, e12)) & (e12 = op1(e12, e10))) | ~ sP35), inference(nnf_transformation, [], [f56])).
fof(f2096, plain, (~ spl144_170 | spl144_62), inference(avatar_split_clause, [], [f662, f1454, f2092])).
fof(f662, plain, ((op1(e10, e10) = e11) | ~ sP36), inference(cnf_transformation, [], [f197])).
fof(f197, plain, ((~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | ~ sP36), inference(nnf_transformation, [], [f57])).
fof(f2088, plain, (~ spl144_169 | spl144_42), inference(avatar_split_clause, [], [f658, f1369, f2084])).
fof(f658, plain, ((e11 = op1(e11, e11)) | ~ sP37), inference(cnf_transformation, [], [f196])).
fof(f196, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | ~ sP37), inference(nnf_transformation, [], [f58])).
fof(f2087, plain, (~ spl144_169 | ~ spl144_42), inference(avatar_split_clause, [], [f659, f1369, f2084])).
fof(f659, plain, (~ (e11 = op1(e11, e11)) | ~ sP37), inference(cnf_transformation, [], [f196])).
fof(f2082, plain, (~ spl144_168 | spl144_27), inference(avatar_split_clause, [], [f652, f1305, f2076])).
fof(f652, plain, ((e12 = op1(e12, e11)) | ~ sP38), inference(cnf_transformation, [], [f195])).
fof(f195, plain, ((~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | ~ sP38), inference(nnf_transformation, [], [f59])).
fof(f2080, plain, (~ spl144_168 | spl144_22), inference(avatar_split_clause, [], [f654, f1284, f2076])).
fof(f654, plain, ((e11 = op1(e12, e12)) | ~ sP38), inference(cnf_transformation, [], [f195])).
fof(f2079, plain, (~ spl144_168 | ~ spl144_27), inference(avatar_split_clause, [], [f655, f1305, f2076])).
fof(f655, plain, (~ (e12 = op1(e12, e11)) | ~ sP38), inference(cnf_transformation, [], [f195])).
fof(f2074, plain, (~ spl144_167 | spl144_27), inference(avatar_split_clause, [], [f648, f1305, f2068])).
fof(f648, plain, ((e12 = op1(e12, e11)) | ~ sP39), inference(cnf_transformation, [], [f194])).
fof(f194, plain, ((~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e12, e12)) & (e12 = op1(e12, e11))) | ~ sP39), inference(nnf_transformation, [], [f60])).
fof(f2073, plain, (~ spl144_167 | ~ spl144_22), inference(avatar_split_clause, [], [f649, f1284, f2068])).
fof(f649, plain, (~ (e11 = op1(e12, e12)) | ~ sP39), inference(cnf_transformation, [], [f194])).
fof(f2071, plain, (~ spl144_167 | ~ spl144_12), inference(avatar_split_clause, [], [f651, f1241, f2068])).
fof(f651, plain, (~ (e13 = op1(e13, e11)) | ~ sP39), inference(cnf_transformation, [], [f194])).
fof(f2064, plain, (~ spl144_166 | spl144_63), inference(avatar_split_clause, [], [f646, f1458, f2060])).
fof(f646, plain, ((op1(e10, e10) = e12) | ~ sP40), inference(cnf_transformation, [], [f193])).
fof(f193, plain, ((~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP40), inference(nnf_transformation, [], [f61])).
fof(f2058, plain, (~ spl144_165 | spl144_23), inference(avatar_split_clause, [], [f640, f1288, f2052])).
fof(f640, plain, ((e12 = op1(e12, e12)) | ~ sP41), inference(cnf_transformation, [], [f192])).
fof(f192, plain, ((~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP41), inference(nnf_transformation, [], [f62])).
fof(f2057, plain, (~ spl144_165 | ~ spl144_23), inference(avatar_split_clause, [], [f641, f1288, f2052])).
fof(f641, plain, (~ (e12 = op1(e12, e12)) | ~ sP41), inference(cnf_transformation, [], [f192])).
fof(f2056, plain, (~ spl144_165 | spl144_43), inference(avatar_split_clause, [], [f642, f1373, f2052])).
fof(f642, plain, ((e12 = op1(e11, e11)) | ~ sP41), inference(cnf_transformation, [], [f192])).
fof(f2050, plain, (~ spl144_164 | spl144_23), inference(avatar_split_clause, [], [f636, f1288, f2044])).
fof(f636, plain, ((e12 = op1(e12, e12)) | ~ sP42), inference(cnf_transformation, [], [f191])).
fof(f191, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP42), inference(nnf_transformation, [], [f63])).
fof(f2049, plain, (~ spl144_164 | ~ spl144_23), inference(avatar_split_clause, [], [f637, f1288, f2044])).
fof(f637, plain, (~ (e12 = op1(e12, e12)) | ~ sP42), inference(cnf_transformation, [], [f191])).
fof(f2042, plain, (~ spl144_163 | spl144_23), inference(avatar_split_clause, [], [f632, f1288, f2036])).
fof(f632, plain, ((e12 = op1(e12, e12)) | ~ sP43), inference(cnf_transformation, [], [f190])).
fof(f190, plain, ((~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP43), inference(nnf_transformation, [], [f64])).
fof(f2041, plain, (~ spl144_163 | ~ spl144_23), inference(avatar_split_clause, [], [f633, f1288, f2036])).
fof(f633, plain, (~ (e12 = op1(e12, e12)) | ~ sP43), inference(cnf_transformation, [], [f190])).
fof(f2034, plain, (~ spl144_162 | spl144_19), inference(avatar_split_clause, [], [f628, f1271, f2028])).
fof(f628, plain, ((e12 = op1(e12, e13)) | ~ sP44), inference(cnf_transformation, [], [f189])).
fof(f189, plain, ((~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | ~ sP44), inference(nnf_transformation, [], [f65])).
fof(f2033, plain, (~ spl144_162 | ~ spl144_24), inference(avatar_split_clause, [], [f629, f1292, f2028])).
fof(f629, plain, (~ (e13 = op1(e12, e12)) | ~ sP44), inference(cnf_transformation, [], [f189])).
fof(f2031, plain, (~ spl144_162 | ~ spl144_49), inference(avatar_split_clause, [], [f631, f1399, f2028])).
fof(f631, plain, (~ (e10 = op1(e10, e13)) | ~ sP44), inference(cnf_transformation, [], [f189])).
fof(f2026, plain, (~ spl144_161 | spl144_19), inference(avatar_split_clause, [], [f624, f1271, f2020])).
fof(f624, plain, ((e12 = op1(e12, e13)) | ~ sP45), inference(cnf_transformation, [], [f188])).
fof(f188, plain, ((~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | ~ sP45), inference(nnf_transformation, [], [f66])).
fof(f2024, plain, (~ spl144_161 | spl144_44), inference(avatar_split_clause, [], [f626, f1377, f2020])).
fof(f626, plain, ((e13 = op1(e11, e11)) | ~ sP45), inference(cnf_transformation, [], [f188])).
fof(f2017, plain, (~ spl144_160 | ~ spl144_24), inference(avatar_split_clause, [], [f621, f1292, f2012])).
fof(f621, plain, (~ (e13 = op1(e12, e12)) | ~ sP46), inference(cnf_transformation, [], [f187])).
fof(f187, plain, ((~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | ~ sP46), inference(nnf_transformation, [], [f67])).
fof(f2016, plain, (~ spl144_160 | spl144_24), inference(avatar_split_clause, [], [f622, f1292, f2012])).
fof(f622, plain, ((e13 = op1(e12, e12)) | ~ sP46), inference(cnf_transformation, [], [f187])).
fof(f2008, plain, (~ spl144_159 | spl144_4), inference(avatar_split_clause, [], [f618, f1207, f2004])).
fof(f618, plain, ((e13 = op1(e13, e13)) | ~ sP47), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13)) & ~ (e13 = op1(e12, e12)) & (e12 = op1(e12, e13))) | ~ sP47), inference(nnf_transformation, [], [f68])).
fof(f2002, plain, (~ spl144_158 | spl144_16), inference(avatar_split_clause, [], [f612, f1258, f1996])).
fof(f612, plain, ((e13 = op1(e13, e10)) | ~ sP48), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | ~ sP48), inference(nnf_transformation, [], [f69])).
fof(f1994, plain, (~ spl144_157 | spl144_16), inference(avatar_split_clause, [], [f608, f1258, f1988])).
fof(f608, plain, ((e13 = op1(e13, e10)) | ~ sP49), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e11 = op1(e11, e10)) & (e10 = op1(e11, e11)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | ~ sP49), inference(nnf_transformation, [], [f70])).
fof(f1986, plain, (~ spl144_156 | spl144_16), inference(avatar_split_clause, [], [f604, f1258, f1980])).
fof(f604, plain, ((e13 = op1(e13, e10)) | ~ sP50), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ((~ (e12 = op1(e12, e10)) & (e10 = op1(e12, e12)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | ~ sP50), inference(nnf_transformation, [], [f71])).
fof(f1978, plain, (~ spl144_155 | spl144_16), inference(avatar_split_clause, [], [f600, f1258, f1972])).
fof(f600, plain, ((e13 = op1(e13, e10)) | ~ sP51), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((~ (e13 = op1(e13, e10)) & (e10 = op1(e13, e13)) & ~ (e10 = op1(e13, e13)) & (e13 = op1(e13, e10))) | ~ sP51), inference(nnf_transformation, [], [f72])).
fof(f1969, plain, (~ spl144_154 | ~ spl144_2), inference(avatar_split_clause, [], [f597, f1199, f1964])).
fof(f597, plain, (~ (e11 = op1(e13, e13)) | ~ sP52), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ((~ (e10 = op1(e10, e11)) & (op1(e10, e10) = e11) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | ~ sP52), inference(nnf_transformation, [], [f73])).
fof(f1961, plain, (~ spl144_153 | ~ spl144_2), inference(avatar_split_clause, [], [f593, f1199, f1956])).
fof(f593, plain, (~ (e11 = op1(e13, e13)) | ~ sP53), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((~ (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | ~ sP53), inference(nnf_transformation, [], [f74])).
fof(f1953, plain, (~ spl144_152 | ~ spl144_2), inference(avatar_split_clause, [], [f589, f1199, f1948])).
fof(f589, plain, (~ (e11 = op1(e13, e13)) | ~ sP54), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((~ (e12 = op1(e12, e11)) & (e11 = op1(e12, e12)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | ~ sP54), inference(nnf_transformation, [], [f75])).
fof(f1945, plain, (~ spl144_151 | ~ spl144_2), inference(avatar_split_clause, [], [f585, f1199, f1940])).
fof(f585, plain, (~ (e11 = op1(e13, e13)) | ~ sP55), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ((~ (e13 = op1(e13, e11)) & (e11 = op1(e13, e13)) & ~ (e11 = op1(e13, e13)) & (e13 = op1(e13, e11))) | ~ sP55), inference(nnf_transformation, [], [f76])).
fof(f1936, plain, (~ spl144_150 | spl144_63), inference(avatar_split_clause, [], [f582, f1458, f1932])).
fof(f582, plain, ((op1(e10, e10) = e12) | ~ sP56), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ((~ (e10 = op1(e10, e12)) & (op1(e10, e10) = e12) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | ~ sP56), inference(nnf_transformation, [], [f77])).
fof(f1930, plain, (~ spl144_149 | spl144_8), inference(avatar_split_clause, [], [f576, f1224, f1924])).
fof(f576, plain, ((e13 = op1(e13, e12)) | ~ sP57), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ((~ (e11 = op1(e11, e12)) & (e12 = op1(e11, e11)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | ~ sP57), inference(nnf_transformation, [], [f78])).
fof(f1928, plain, (~ spl144_149 | spl144_43), inference(avatar_split_clause, [], [f578, f1373, f1924])).
fof(f578, plain, ((e12 = op1(e11, e11)) | ~ sP57), inference(cnf_transformation, [], [f176])).
fof(f1927, plain, (~ spl144_149 | ~ spl144_38), inference(avatar_split_clause, [], [f579, f1352, f1924])).
fof(f579, plain, (~ (e11 = op1(e11, e12)) | ~ sP57), inference(cnf_transformation, [], [f176])).
fof(f1922, plain, (~ spl144_148 | spl144_8), inference(avatar_split_clause, [], [f572, f1224, f1916])).
fof(f572, plain, ((e13 = op1(e13, e12)) | ~ sP58), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | ~ sP58), inference(nnf_transformation, [], [f79])).
fof(f1920, plain, (~ spl144_148 | spl144_23), inference(avatar_split_clause, [], [f574, f1288, f1916])).
fof(f574, plain, ((e12 = op1(e12, e12)) | ~ sP58), inference(cnf_transformation, [], [f175])).
fof(f1919, plain, (~ spl144_148 | ~ spl144_23), inference(avatar_split_clause, [], [f575, f1288, f1916])).
fof(f575, plain, (~ (e12 = op1(e12, e12)) | ~ sP58), inference(cnf_transformation, [], [f175])).
fof(f1914, plain, (~ spl144_147 | spl144_8), inference(avatar_split_clause, [], [f568, f1224, f1908])).
fof(f568, plain, ((e13 = op1(e13, e12)) | ~ sP59), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e13 = op1(e13, e12)) & (e12 = op1(e13, e13)) & ~ (e12 = op1(e13, e13)) & (e13 = op1(e13, e12))) | ~ sP59), inference(nnf_transformation, [], [f80])).
fof(f1913, plain, (~ spl144_147 | ~ spl144_3), inference(avatar_split_clause, [], [f569, f1203, f1908])).
fof(f569, plain, (~ (e12 = op1(e13, e13)) | ~ sP59), inference(cnf_transformation, [], [f174])).
fof(f1912, plain, (~ spl144_147 | spl144_3), inference(avatar_split_clause, [], [f570, f1203, f1908])).
fof(f570, plain, ((e12 = op1(e13, e13)) | ~ sP59), inference(cnf_transformation, [], [f174])).
fof(f1906, plain, (~ spl144_146 | spl144_4), inference(avatar_split_clause, [], [f564, f1207, f1900])).
fof(f564, plain, ((e13 = op1(e13, e13)) | ~ sP60), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e10 = op1(e10, e13)) & (op1(e10, e10) = e13) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP60), inference(nnf_transformation, [], [f81])).
fof(f1898, plain, (~ spl144_145 | spl144_4), inference(avatar_split_clause, [], [f560, f1207, f1892])).
fof(f560, plain, ((e13 = op1(e13, e13)) | ~ sP61), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e11 = op1(e11, e13)) & (e13 = op1(e11, e11)) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP61), inference(nnf_transformation, [], [f82])).
fof(f1890, plain, (~ spl144_144 | spl144_4), inference(avatar_split_clause, [], [f556, f1207, f1884])).
fof(f556, plain, ((e13 = op1(e13, e13)) | ~ sP62), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e12 = op1(e12, e13)) & (e13 = op1(e12, e12)) & ~ (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ~ sP62), inference(nnf_transformation, [], [f83])).
fof(f1882, plain, (~ spl144_139 | spl144_143), inference(avatar_split_clause, [], [f548, f1879, f1857])).
fof(f548, plain, ((e10 = op1(e10, op1(e10, e10))) | ~ sP63), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, op1(e10, e13))) & (e12 = op1(e10, op1(e10, e12))) & (e11 = op1(e10, op1(e10, e11))) & (e10 = op1(e10, op1(e10, e10)))) | ~ sP63), inference(nnf_transformation, [], [f84])).
fof(f1861, plain, (~ spl144_139 | ~ spl144_21 | spl144_31), inference(avatar_split_clause, [], [f554, f1322, f1280, f1857])).
fof(f554, plain, ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12)) | ~ sP63), inference(cnf_transformation, [], [f170])).
fof(f1833, plain, (~ spl144_134 | ~ spl144_2 | spl144_12), inference(avatar_split_clause, [], [f547, f1241, f1199, f1830])).
fof(f547, plain, ((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13)) | ~ sP64), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e13 = op1(e11, op1(e11, e13))) & (e12 = op1(e11, op1(e11, e12))) & (e11 = op1(e11, op1(e11, e11))) & (e10 = op1(e11, op1(e11, e10)))) | ~ sP64), inference(nnf_transformation, [], [f85])).
fof(f1828, plain, (~ spl144_129 | spl144_133), inference(avatar_split_clause, [], [f532, f1825, f1803])).
fof(f532, plain, ((e10 = op1(e12, op1(e12, e10))) | ~ sP65), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e13 = op1(e12, op1(e12, e13))) & (e12 = op1(e12, op1(e12, e12))) & (e11 = op1(e12, op1(e12, e11))) & (e10 = op1(e12, op1(e12, e10)))) | ~ sP65), inference(nnf_transformation, [], [f86])).
fof(f1801, plain, (spl144_125 | spl144_121 | spl144_117 | spl144_113), inference(avatar_split_clause, [], [f376, f1703, f1720, f1737, f1754])).
fof(f376, plain, ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax4)).
fof(f1800, plain, (spl144_125 | spl144_109 | spl144_93 | spl144_77), inference(avatar_split_clause, [], [f377, f1550, f1618, f1686, f1754])).
fof(f377, plain, ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f1799, plain, (spl144_126 | spl144_122 | spl144_118 | spl144_114), inference(avatar_split_clause, [], [f378, f1707, f1724, f1741, f1758])).
fof(f378, plain, ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)), inference(cnf_transformation, [], [f4])).
fof(f1798, plain, (spl144_126 | spl144_110 | spl144_94 | spl144_78), inference(avatar_split_clause, [], [f379, f1554, f1622, f1690, f1758])).
fof(f379, plain, ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)), inference(cnf_transformation, [], [f4])).
fof(f1791, plain, (spl144_110 | spl144_106 | spl144_102 | spl144_98), inference(avatar_split_clause, [], [f386, f1639, f1656, f1673, f1690])).
fof(f386, plain, ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1788, plain, (spl144_123 | spl144_107 | spl144_91 | spl144_75), inference(avatar_split_clause, [], [f389, f1541, f1609, f1677, f1745])).
fof(f389, plain, ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1787, plain, (spl144_112 | spl144_108 | spl144_104 | spl144_100), inference(avatar_split_clause, [], [f390, f1647, f1664, f1681, f1698])).
fof(f390, plain, ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1786, plain, (spl144_124 | spl144_108 | spl144_92 | spl144_76), inference(avatar_split_clause, [], [f391, f1545, f1613, f1681, f1749])).
fof(f391, plain, ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1785, plain, (spl144_93 | spl144_89 | spl144_85 | spl144_81), inference(avatar_split_clause, [], [f392, f1567, f1584, f1601, f1618])).
fof(f392, plain, ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1784, plain, (spl144_117 | spl144_101 | spl144_85 | spl144_69), inference(avatar_split_clause, [], [f393, f1516, f1584, f1652, f1720])).
fof(f393, plain, ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1783, plain, (spl144_94 | spl144_90 | spl144_86 | spl144_82), inference(avatar_split_clause, [], [f394, f1571, f1588, f1605, f1622])).
fof(f394, plain, ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1782, plain, (spl144_118 | spl144_102 | spl144_86 | spl144_70), inference(avatar_split_clause, [], [f395, f1520, f1588, f1656, f1724])).
fof(f395, plain, ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1780, plain, (spl144_119 | spl144_103 | spl144_87 | spl144_71), inference(avatar_split_clause, [], [f397, f1524, f1592, f1660, f1728])).
fof(f397, plain, ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1779, plain, (spl144_96 | spl144_92 | spl144_88 | spl144_84), inference(avatar_split_clause, [], [f398, f1579, f1596, f1613, f1630])).
fof(f398, plain, ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1777, plain, (spl144_77 | spl144_73 | spl144_69 | spl144_65), inference(avatar_split_clause, [], [f400, f1499, f1516, f1533, f1550])).
fof(f400, plain, ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1772, plain, (spl144_115 | spl144_99 | spl144_83 | spl144_67), inference(avatar_split_clause, [], [f405, f1507, f1575, f1643, f1711])).
fof(f405, plain, ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1771, plain, (spl144_80 | spl144_76 | spl144_72 | spl144_68), inference(avatar_split_clause, [], [f406, f1511, f1528, f1545, f1562])).
fof(f406, plain, ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1752, plain, (spl144_121 | spl144_122 | spl144_123 | spl144_124), inference(avatar_split_clause, [], [f361, f1749, f1745, f1741, f1737])).
fof(f361, plain, ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax3)).
fof(f1735, plain, (spl144_117 | spl144_118 | spl144_119 | spl144_120), inference(avatar_split_clause, [], [f362, f1732, f1728, f1724, f1720])).
fof(f362, plain, ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f3])).
fof(f1718, plain, (spl144_113 | spl144_114 | spl144_115 | spl144_116), inference(avatar_split_clause, [], [f363, f1715, f1711, f1707, f1703])).
fof(f363, plain, ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f3])).
fof(f1701, plain, (spl144_109 | spl144_110 | spl144_111 | spl144_112), inference(avatar_split_clause, [], [f364, f1698, f1694, f1690, f1686])).
fof(f364, plain, ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f3])).
fof(f1667, plain, (spl144_101 | spl144_102 | spl144_103 | spl144_104), inference(avatar_split_clause, [], [f366, f1664, f1660, f1656, f1652])).
fof(f366, plain, ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))), inference(cnf_transformation, [], [f3])).
fof(f1633, plain, (spl144_93 | spl144_94 | spl144_95 | spl144_96), inference(avatar_split_clause, [], [f368, f1630, f1626, f1622, f1618])).
fof(f368, plain, ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f3])).
fof(f1616, plain, (spl144_89 | spl144_90 | spl144_91 | spl144_92), inference(avatar_split_clause, [], [f369, f1613, f1609, f1605, f1601])).
fof(f369, plain, ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))), inference(cnf_transformation, [], [f3])).
fof(f1599, plain, (spl144_85 | spl144_86 | spl144_87 | spl144_88), inference(avatar_split_clause, [], [f370, f1596, f1592, f1588, f1584])).
fof(f370, plain, ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))), inference(cnf_transformation, [], [f3])).
fof(f1582, plain, (spl144_81 | spl144_82 | spl144_83 | spl144_84), inference(avatar_split_clause, [], [f371, f1579, f1575, f1571, f1567])).
fof(f371, plain, ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))), inference(cnf_transformation, [], [f3])).
fof(f1548, plain, (spl144_73 | spl144_74 | spl144_75 | spl144_76), inference(avatar_split_clause, [], [f373, f1545, f1541, f1537, f1533])).
fof(f373, plain, ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))), inference(cnf_transformation, [], [f3])).
fof(f1531, plain, (spl144_69 | spl144_70 | spl144_71 | spl144_72), inference(avatar_split_clause, [], [f374, f1528, f1524, f1520, f1516])).
fof(f374, plain, ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))), inference(cnf_transformation, [], [f3])).
fof(f1495, plain, (spl144_62 | spl144_58 | spl144_54 | spl144_50), inference(avatar_split_clause, [], [f330, f1403, f1420, f1437, f1454])).
fof(f330, plain, ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax2)).
fof(f1494, plain, (spl144_62 | spl144_46 | spl144_30 | spl144_14), inference(avatar_split_clause, [], [f331, f1250, f1318, f1386, f1454])).
fof(f331, plain, ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)), inference(cnf_transformation, [], [f2])).
fof(f1493, plain, (spl144_63 | spl144_59 | spl144_55 | spl144_51), inference(avatar_split_clause, [], [f332, f1407, f1424, f1441, f1458])).
fof(f332, plain, ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)), inference(cnf_transformation, [], [f2])).
fof(f1489, plain, (spl144_45 | spl144_41 | spl144_37 | spl144_33), inference(avatar_split_clause, [], [f336, f1331, f1348, f1365, f1382])).
fof(f336, plain, ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1487, plain, (spl144_46 | spl144_42 | spl144_38 | spl144_34), inference(avatar_split_clause, [], [f338, f1335, f1352, f1369, f1386])).
fof(f338, plain, ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1486, plain, (spl144_58 | spl144_42 | spl144_26 | spl144_10), inference(avatar_split_clause, [], [f339, f1233, f1301, f1369, f1437])).
fof(f339, plain, ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f1485, plain, (spl144_47 | spl144_43 | spl144_39 | spl144_35), inference(avatar_split_clause, [], [f340, f1339, f1356, f1373, f1390])).
fof(f340, plain, ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1484, plain, (spl144_59 | spl144_43 | spl144_27 | spl144_11), inference(avatar_split_clause, [], [f341, f1237, f1305, f1373, f1441])).
fof(f341, plain, ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f1480, plain, (spl144_53 | spl144_37 | spl144_21 | spl144_5), inference(avatar_split_clause, [], [f345, f1212, f1280, f1348, f1416])).
fof(f345, plain, ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1479, plain, (spl144_30 | spl144_26 | spl144_22 | spl144_18), inference(avatar_split_clause, [], [f346, f1267, f1284, f1301, f1318])).
fof(f346, plain, ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1478, plain, (spl144_54 | spl144_38 | spl144_22 | spl144_6), inference(avatar_split_clause, [], [f347, f1216, f1284, f1352, f1420])).
fof(f347, plain, ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1477, plain, (spl144_31 | spl144_27 | spl144_23 | spl144_19), inference(avatar_split_clause, [], [f348, f1271, f1288, f1305, f1322])).
fof(f348, plain, ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1467, plain, (spl144_16 | spl144_12 | spl144_8 | spl144_4), inference(avatar_split_clause, [], [f358, f1207, f1224, f1241, f1258])).
fof(f358, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1466, plain, (spl144_52 | spl144_36 | spl144_20 | spl144_4), inference(avatar_split_clause, [], [f359, f1207, f1275, f1343, f1411])).
fof(f359, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1448, plain, (spl144_57 | spl144_58 | spl144_59 | spl144_60), inference(avatar_split_clause, [], [f313, f1445, f1441, f1437, f1433])).
fof(f313, plain, ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG129+1.p', ax1)).
fof(f1431, plain, (spl144_53 | spl144_54 | spl144_55 | spl144_56), inference(avatar_split_clause, [], [f314, f1428, f1424, f1420, f1416])).
fof(f314, plain, ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f1])).
fof(f1414, plain, (spl144_49 | spl144_50 | spl144_51 | spl144_52), inference(avatar_split_clause, [], [f315, f1411, f1407, f1403, f1399])).
fof(f315, plain, ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f1])).
fof(f1397, plain, (spl144_45 | spl144_46 | spl144_47 | spl144_48), inference(avatar_split_clause, [], [f316, f1394, f1390, f1386, f1382])).
fof(f316, plain, ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f1])).
fof(f1329, plain, (spl144_29 | spl144_30 | spl144_31 | spl144_32), inference(avatar_split_clause, [], [f320, f1326, f1322, f1318, f1314])).
fof(f320, plain, ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f1])).
fof(f1312, plain, (spl144_25 | spl144_26 | spl144_27 | spl144_28), inference(avatar_split_clause, [], [f321, f1309, f1305, f1301, f1297])).
fof(f321, plain, ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f1295, plain, (spl144_21 | spl144_22 | spl144_23 | spl144_24), inference(avatar_split_clause, [], [f322, f1292, f1288, f1284, f1280])).
fof(f322, plain, ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))), inference(cnf_transformation, [], [f1])).
fof(f1244, plain, (spl144_9 | spl144_10 | spl144_11 | spl144_12), inference(avatar_split_clause, [], [f325, f1241, f1237, f1233, f1229])).
fof(f325, plain, ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))), inference(cnf_transformation, [], [f1])).
fof(f1227, plain, (spl144_5 | spl144_6 | spl144_7 | spl144_8), inference(avatar_split_clause, [], [f326, f1224, f1220, f1216, f1212])).
fof(f326, plain, ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))), inference(cnf_transformation, [], [f1])).