fof(f15372, plain, $false, inference(avatar_sat_refutation, [], [f1573, f1594, f1826, f1828, f1837, f1838, f1842, f1863, f1865, f1874, f2274, f2316, f2400, f2402, f2408, f2413, f2416, f2426, f2427, f2441, f2452, f2454, f2940, f2967, f2994, f3005, f5136, f8314, f8423, f8459, f8524, f8545, f8555, f8565, f8589, f8784, f8797, f8813, f8835, f8873, f8899, f8942, f8993, f9153, f9186, f9209, f9210, f9242, f9245, f9253, f9365, f9369, f9380, f9383, f9384, f9385, f9392, f9420, f9426, f9433, f9447, f9471, f9500, f9511, f9555, f9646, f9659, f9660, f9726, f9844, f9845, f9850, f9862, f9892, f9914, f9924, f9928, f9948, f9958, f9990, f9992, f9995, f10002, f10013, f10018, f10019, f10035, f10098, f10104, f10158, f10160, f10254, f10325, f10346, f10380, f10383, f10417, f10422, f10453, f10514, f10515, f10529, f10530, f10537, f10551, f10552, f10614, f10623, f10631, f10759, f10772, f10857, f10907, f10915, f10917, f10922, f10952, f10958, f10968, f10972, f10980, f11033, f11065, f11076, f11078, f11080, f11093, f11098, f11137, f11172, f11200, f11229, f11273, f11290, f11309, f11360, f11364, f11438, f11452, f11453, f11468, f11482, f11489, f11509, f11542, f11551, f11577, f11580, f11669, f11729, f11761, f11792, f11798, f11813, f11815, f11816, f11858, f11878, f11905, f11913, f11930, f12031, f12034, f12185, f12367, f12545, f12717, f12892, f13069, f13242, f13405, f13560, f13711, f13870, f14006, f14142, f14269, f14385, f14535, f14684, f14823, f14960, f15095, f15231, f15352])).
fof(f15352, plain, (~ spl80_119 | spl80_666), inference(avatar_contradiction_clause, [], [f15351])).
fof(f15351, plain, ($false | (~ spl80_119 | spl80_666)), inference(subsumption_resolution, [], [f15350, f782])).
fof(f782, plain, (e20 = h17(e13)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ((h17(e14) = op2(op2(e20, e24), op2(e24, op2(e20, e24)))) & (h17(e12) = op2(e24, op2(e20, e24))) & (op2(e20, e24) = h17(e10)) & (e20 = h17(e13)) & (e24 = h17(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax32)).
fof(f15350, plain, (~ (e20 = h17(e13)) | (~ spl80_119 | spl80_666)), inference(forward_demodulation, [], [f5081, f1799])).
fof(f1799, plain, ((e13 = op1(e10, e11)) | ~ spl80_119), inference(avatar_component_clause, [], [f1797])).
fof(f1797, plain, (spl80_119 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_119])])).
fof(f5081, plain, (~ (e20 = h17(op1(e10, e11))) | spl80_666), inference(avatar_component_clause, [], [f5079])).
fof(f5079, plain, (spl80_666 <=> (e20 = h17(op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_666])])).
fof(f15231, plain, (~ spl80_115 | spl80_665), inference(avatar_contradiction_clause, [], [f15230])).
fof(f15230, plain, ($false | (~ spl80_115 | spl80_665)), inference(trivial_inequality_removal, [], [f15229])).
fof(f15229, plain, (~ (h17(e14) = h17(e14)) | (~ spl80_115 | spl80_665)), inference(forward_demodulation, [], [f5077, f1782])).
fof(f1782, plain, ((e14 = op1(e10, e12)) | ~ spl80_115), inference(avatar_component_clause, [], [f1780])).
fof(f1780, plain, (spl80_115 <=> (e14 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_115])])).
fof(f5077, plain, (~ (h17(e14) = h17(op1(e10, e12))) | spl80_665), inference(avatar_component_clause, [], [f5075])).
fof(f5075, plain, (spl80_665 <=> (h17(e14) = h17(op1(e10, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl80_665])])).
fof(f15095, plain, (~ spl80_89 | ~ spl80_136 | ~ spl80_305 | spl80_661), inference(avatar_contradiction_clause, [], [f15094])).
fof(f15094, plain, ($false | (~ spl80_89 | ~ spl80_136 | ~ spl80_305 | spl80_661)), inference(subsumption_resolution, [], [f15093, f1921])).
fof(f1921, plain, ((e20 = op2(e24, e22)) | ~ spl80_136), inference(avatar_component_clause, [], [f1919])).
fof(f1919, plain, (spl80_136 <=> (e20 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_136])])).
fof(f15093, plain, (~ (e20 = op2(e24, e22)) | (~ spl80_89 | ~ spl80_305 | spl80_661)), inference(forward_demodulation, [], [f15092, f782])).
fof(f15092, plain, (~ (op2(e24, e22) = h17(e13)) | (~ spl80_89 | ~ spl80_305 | spl80_661)), inference(forward_demodulation, [], [f15091, f1673])).
fof(f1673, plain, ((e13 = op1(e11, e12)) | ~ spl80_89), inference(avatar_component_clause, [], [f1671])).
fof(f1671, plain, (spl80_89 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_89])])).
fof(f15091, plain, (~ (op2(e24, e22) = h17(op1(e11, e12))) | (~ spl80_305 | spl80_661)), inference(forward_demodulation, [], [f5061, f2965])).
fof(f2965, plain, ((e22 = h17(e12)) | ~ spl80_305), inference(avatar_component_clause, [], [f2964])).
fof(f2964, plain, (spl80_305 <=> (e22 = h17(e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_305])])).
fof(f5061, plain, (~ (h17(op1(e11, e12)) = op2(e24, h17(e12))) | spl80_661), inference(avatar_component_clause, [], [f5059])).
fof(f5059, plain, (spl80_661 <=> (h17(op1(e11, e12)) = op2(e24, h17(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl80_661])])).
fof(f14960, plain, (~ spl80_85 | ~ spl80_300 | ~ spl80_496 | spl80_660), inference(avatar_contradiction_clause, [], [f14959])).
fof(f14959, plain, ($false | (~ spl80_85 | ~ spl80_300 | ~ spl80_496 | spl80_660)), inference(subsumption_resolution, [], [f14958, f4000])).
fof(f4000, plain, ((e23 = h4(e10)) | ~ spl80_496), inference(avatar_component_clause, [], [f3999])).
fof(f3999, plain, (spl80_496 <=> (e23 = h4(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_496])])).
fof(f14958, plain, (~ (e23 = h4(e10)) | (~ spl80_85 | ~ spl80_300 | spl80_660)), inference(forward_demodulation, [], [f14957, f2938])).
fof(f2938, plain, ((e23 = h17(e14)) | ~ spl80_300), inference(avatar_component_clause, [], [f2937])).
fof(f2937, plain, (spl80_300 <=> (e23 = h17(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_300])])).
fof(f14957, plain, (~ (h4(e10) = h17(e14)) | (~ spl80_85 | spl80_660)), inference(forward_demodulation, [], [f5057, f1656])).
fof(f1656, plain, ((e14 = op1(e11, e13)) | ~ spl80_85), inference(avatar_component_clause, [], [f1654])).
fof(f1654, plain, (spl80_85 <=> (e14 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_85])])).
fof(f5057, plain, (~ (h4(e10) = h17(op1(e11, e13))) | spl80_660), inference(avatar_component_clause, [], [f5055])).
fof(f5055, plain, (spl80_660 <=> (h4(e10) = h17(op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_660])])).
fof(f14823, plain, (~ spl80_98 | spl80_663), inference(avatar_contradiction_clause, [], [f14822])).
fof(f14822, plain, ($false | (~ spl80_98 | spl80_663)), inference(trivial_inequality_removal, [], [f14821])).
fof(f14821, plain, (~ (h17(e12) = h17(e12)) | (~ spl80_98 | spl80_663)), inference(forward_demodulation, [], [f5069, f1711])).
fof(f1711, plain, ((e12 = op1(e11, e10)) | ~ spl80_98), inference(avatar_component_clause, [], [f1709])).
fof(f1709, plain, (spl80_98 <=> (e12 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_98])])).
fof(f5069, plain, (~ (h17(e12) = h17(op1(e11, e10))) | spl80_663), inference(avatar_component_clause, [], [f5067])).
fof(f5067, plain, (spl80_663 <=> (h17(e12) = h17(op1(e11, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_663])])).
fof(f14684, plain, (~ spl80_70 | ~ spl80_179 | ~ spl80_300 | ~ spl80_305 | spl80_658), inference(avatar_contradiction_clause, [], [f14683])).
fof(f14683, plain, ($false | (~ spl80_70 | ~ spl80_179 | ~ spl80_300 | ~ spl80_305 | spl80_658)), inference(subsumption_resolution, [], [f14682, f2101])).
fof(f2101, plain, ((e23 = op2(e22, e24)) | ~ spl80_179), inference(avatar_component_clause, [], [f2099])).
fof(f2099, plain, (spl80_179 <=> (e23 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_179])])).
fof(f14682, plain, (~ (e23 = op2(e22, e24)) | (~ spl80_70 | ~ spl80_300 | ~ spl80_305 | spl80_658)), inference(forward_demodulation, [], [f14681, f2938])).
fof(f14681, plain, (~ (op2(e22, e24) = h17(e14)) | (~ spl80_70 | ~ spl80_305 | spl80_658)), inference(forward_demodulation, [], [f14680, f1593])).
fof(f1593, plain, ((e14 = op1(e12, e11)) | ~ spl80_70), inference(avatar_component_clause, [], [f1591])).
fof(f1591, plain, (spl80_70 <=> (e14 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_70])])).
fof(f14680, plain, (~ (op2(e22, e24) = h17(op1(e12, e11))) | (~ spl80_305 | spl80_658)), inference(forward_demodulation, [], [f5049, f2965])).
fof(f5049, plain, (~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | spl80_658), inference(avatar_component_clause, [], [f5047])).
fof(f5047, plain, (spl80_658 <=> (h17(op1(e12, e11)) = op2(h17(e12), e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_658])])).
fof(f14535, plain, (~ spl80_50 | ~ spl80_244 | ~ spl80_300 | ~ spl80_310 | spl80_656), inference(avatar_contradiction_clause, [], [f14534])).
fof(f14534, plain, ($false | (~ spl80_50 | ~ spl80_244 | ~ spl80_300 | ~ spl80_310 | spl80_656)), inference(subsumption_resolution, [], [f14533, f2374])).
fof(f2374, plain, ((e23 = op2(e20, e21)) | ~ spl80_244), inference(avatar_component_clause, [], [f2372])).
fof(f2372, plain, (spl80_244 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_244])])).
fof(f14533, plain, (~ (e23 = op2(e20, e21)) | (~ spl80_50 | ~ spl80_300 | ~ spl80_310 | spl80_656)), inference(forward_demodulation, [], [f14532, f2938])).
fof(f14532, plain, (~ (op2(e20, e21) = h17(e14)) | (~ spl80_50 | ~ spl80_310 | spl80_656)), inference(forward_demodulation, [], [f14531, f1509])).
fof(f1509, plain, ((e14 = op1(e13, e10)) | ~ spl80_50), inference(avatar_component_clause, [], [f1507])).
fof(f1507, plain, (spl80_50 <=> (e14 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_50])])).
fof(f14531, plain, (~ (op2(e20, e21) = h17(op1(e13, e10))) | (~ spl80_310 | spl80_656)), inference(forward_demodulation, [], [f5041, f2992])).
fof(f2992, plain, ((e21 = h17(e10)) | ~ spl80_310), inference(avatar_component_clause, [], [f2991])).
fof(f2991, plain, (spl80_310 <=> (e21 = h17(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_310])])).
fof(f5041, plain, (~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | spl80_656), inference(avatar_component_clause, [], [f5039])).
fof(f5039, plain, (spl80_656 <=> (h17(op1(e13, e10)) = op2(e20, h17(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_656])])).
fof(f14385, plain, (~ spl80_37 | ~ spl80_240 | ~ spl80_305 | spl80_654), inference(avatar_contradiction_clause, [], [f14384])).
fof(f14384, plain, ($false | (~ spl80_37 | ~ spl80_240 | ~ spl80_305 | spl80_654)), inference(subsumption_resolution, [], [f14383, f2357])).
fof(f2357, plain, ((e24 = op2(e20, e22)) | ~ spl80_240), inference(avatar_component_clause, [], [f2355])).
fof(f2355, plain, (spl80_240 <=> (e24 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_240])])).
fof(f14383, plain, (~ (e24 = op2(e20, e22)) | (~ spl80_37 | ~ spl80_305 | spl80_654)), inference(forward_demodulation, [], [f14382, f781])).
fof(f781, plain, (e24 = h17(e11)), inference(cnf_transformation, [], [f32])).
fof(f14382, plain, (~ (op2(e20, e22) = h17(e11)) | (~ spl80_37 | ~ spl80_305 | spl80_654)), inference(forward_demodulation, [], [f14381, f1455])).
fof(f1455, plain, ((e11 = op1(e13, e12)) | ~ spl80_37), inference(avatar_component_clause, [], [f1453])).
fof(f1453, plain, (spl80_37 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_37])])).
fof(f14381, plain, (~ (op2(e20, e22) = h17(op1(e13, e12))) | (~ spl80_305 | spl80_654)), inference(forward_demodulation, [], [f5033, f2965])).
fof(f5033, plain, (~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | spl80_654), inference(avatar_component_clause, [], [f5031])).
fof(f5031, plain, (spl80_654 <=> (h17(op1(e13, e12)) = op2(e20, h17(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl80_654])])).
fof(f14269, plain, (~ spl80_28 | ~ spl80_233 | ~ spl80_300 | ~ spl80_305 | spl80_652), inference(avatar_contradiction_clause, [], [f14268])).
fof(f14268, plain, ($false | (~ spl80_28 | ~ spl80_233 | ~ spl80_300 | ~ spl80_305 | spl80_652)), inference(subsumption_resolution, [], [f14267, f2328])).
fof(f2328, plain, ((e22 = op2(e20, e23)) | ~ spl80_233), inference(avatar_component_clause, [], [f2326])).
fof(f2326, plain, (spl80_233 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_233])])).
fof(f14267, plain, (~ (e22 = op2(e20, e23)) | (~ spl80_28 | ~ spl80_300 | ~ spl80_305 | spl80_652)), inference(forward_demodulation, [], [f14266, f2965])).
fof(f14266, plain, (~ (op2(e20, e23) = h17(e12)) | (~ spl80_28 | ~ spl80_300 | spl80_652)), inference(forward_demodulation, [], [f14265, f1417])).
fof(f1417, plain, ((e12 = op1(e13, e14)) | ~ spl80_28), inference(avatar_component_clause, [], [f1415])).
fof(f1415, plain, (spl80_28 <=> (e12 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_28])])).
fof(f14265, plain, (~ (op2(e20, e23) = h17(op1(e13, e14))) | (~ spl80_300 | spl80_652)), inference(forward_demodulation, [], [f5025, f2938])).
fof(f5025, plain, (~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | spl80_652), inference(avatar_component_clause, [], [f5023])).
fof(f5023, plain, (spl80_652 <=> (h17(op1(e13, e14)) = op2(e20, h17(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_652])])).
fof(f14142, plain, (~ spl80_18 | ~ spl80_153 | ~ spl80_300 | ~ spl80_305 | spl80_651), inference(avatar_contradiction_clause, [], [f14141])).
fof(f14141, plain, ($false | (~ spl80_18 | ~ spl80_153 | ~ spl80_300 | ~ spl80_305 | spl80_651)), inference(subsumption_resolution, [], [f14140, f1992])).
fof(f1992, plain, ((e22 = op2(e23, e24)) | ~ spl80_153), inference(avatar_component_clause, [], [f1990])).
fof(f1990, plain, (spl80_153 <=> (e22 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_153])])).
fof(f14140, plain, (~ (e22 = op2(e23, e24)) | (~ spl80_18 | ~ spl80_300 | ~ spl80_305 | spl80_651)), inference(forward_demodulation, [], [f14139, f2965])).
fof(f14139, plain, (~ (op2(e23, e24) = h17(e12)) | (~ spl80_18 | ~ spl80_300 | spl80_651)), inference(forward_demodulation, [], [f14138, f1375])).
fof(f1375, plain, ((e12 = op1(e14, e11)) | ~ spl80_18), inference(avatar_component_clause, [], [f1373])).
fof(f1373, plain, (spl80_18 <=> (e12 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_18])])).
fof(f14138, plain, (~ (op2(e23, e24) = h17(op1(e14, e11))) | (~ spl80_300 | spl80_651)), inference(forward_demodulation, [], [f5021, f2938])).
fof(f5021, plain, (~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | spl80_651), inference(avatar_component_clause, [], [f5019])).
fof(f5019, plain, (spl80_651 <=> (h17(op1(e14, e11)) = op2(h17(e14), e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_651])])).
fof(f14006, plain, (~ spl80_5 | ~ spl80_159 | ~ spl80_300 | spl80_648), inference(avatar_contradiction_clause, [], [f14005])).
fof(f14005, plain, ($false | (~ spl80_5 | ~ spl80_159 | ~ spl80_300 | spl80_648)), inference(subsumption_resolution, [], [f14004, f2017])).
fof(f2017, plain, ((e23 = op2(e23, e23)) | ~ spl80_159), inference(avatar_component_clause, [], [f2015])).
fof(f2015, plain, (spl80_159 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_159])])).
fof(f14004, plain, (~ (e23 = op2(e23, e23)) | (~ spl80_5 | ~ spl80_300 | spl80_648)), inference(forward_demodulation, [], [f14003, f2938])).
fof(f14003, plain, (~ (op2(e23, e23) = h17(e14)) | (~ spl80_5 | ~ spl80_300 | spl80_648)), inference(forward_demodulation, [], [f14002, f1320])).
fof(f1320, plain, ((e14 = op1(e14, e14)) | ~ spl80_5), inference(avatar_component_clause, [], [f1318])).
fof(f1318, plain, (spl80_5 <=> (e14 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_5])])).
fof(f14002, plain, (~ (op2(e23, e23) = h17(op1(e14, e14))) | (~ spl80_300 | spl80_648)), inference(forward_demodulation, [], [f5009, f2938])).
fof(f5009, plain, (~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | spl80_648), inference(avatar_component_clause, [], [f5007])).
fof(f5007, plain, (spl80_648 <=> (h17(op1(e14, e14)) = op2(h17(e14), h17(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_648])])).
fof(f13870, plain, (~ spl80_11 | ~ spl80_162 | ~ spl80_300 | ~ spl80_305 | ~ spl80_310 | spl80_647), inference(avatar_contradiction_clause, [], [f13869])).
fof(f13869, plain, ($false | (~ spl80_11 | ~ spl80_162 | ~ spl80_300 | ~ spl80_305 | ~ spl80_310 | spl80_647)), inference(subsumption_resolution, [], [f13868, f2030])).
fof(f2030, plain, ((e21 = op2(e23, e22)) | ~ spl80_162), inference(avatar_component_clause, [], [f2028])).
fof(f2028, plain, (spl80_162 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_162])])).
fof(f13868, plain, (~ (e21 = op2(e23, e22)) | (~ spl80_11 | ~ spl80_300 | ~ spl80_305 | ~ spl80_310 | spl80_647)), inference(forward_demodulation, [], [f13867, f2992])).
fof(f13867, plain, (~ (op2(e23, e22) = h17(e10)) | (~ spl80_11 | ~ spl80_300 | ~ spl80_305 | spl80_647)), inference(forward_demodulation, [], [f13866, f1346])).
fof(f1346, plain, ((e10 = op1(e14, e12)) | ~ spl80_11), inference(avatar_component_clause, [], [f1344])).
fof(f1344, plain, (spl80_11 <=> (e10 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_11])])).
fof(f13866, plain, (~ (op2(e23, e22) = h17(op1(e14, e12))) | (~ spl80_300 | ~ spl80_305 | spl80_647)), inference(forward_demodulation, [], [f13865, f2938])).
fof(f13865, plain, (~ (h17(op1(e14, e12)) = op2(h17(e14), e22)) | (~ spl80_305 | spl80_647)), inference(forward_demodulation, [], [f5005, f2965])).
fof(f5005, plain, (~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | spl80_647), inference(avatar_component_clause, [], [f5003])).
fof(f5003, plain, (spl80_647 <=> (h17(op1(e14, e12)) = op2(h17(e14), h17(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl80_647])])).
fof(f13711, plain, (~ spl80_108 | ~ spl80_223 | ~ spl80_305 | ~ spl80_310 | spl80_664), inference(avatar_contradiction_clause, [], [f13710])).
fof(f13710, plain, ($false | (~ spl80_108 | ~ spl80_223 | ~ spl80_305 | ~ spl80_310 | spl80_664)), inference(subsumption_resolution, [], [f13709, f2286])).
fof(f2286, plain, ((e22 = op2(e21, e20)) | ~ spl80_223), inference(avatar_component_clause, [], [f2284])).
fof(f2284, plain, (spl80_223 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_223])])).
fof(f13709, plain, (~ (e22 = op2(e21, e20)) | (~ spl80_108 | ~ spl80_305 | ~ spl80_310 | spl80_664)), inference(forward_demodulation, [], [f13708, f2965])).
fof(f13708, plain, (~ (op2(e21, e20) = h17(e12)) | (~ spl80_108 | ~ spl80_310 | spl80_664)), inference(forward_demodulation, [], [f13707, f1753])).
fof(f1753, plain, ((e12 = op1(e10, e13)) | ~ spl80_108), inference(avatar_component_clause, [], [f1751])).
fof(f1751, plain, (spl80_108 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_108])])).
fof(f13707, plain, (~ (op2(e21, e20) = h17(op1(e10, e13))) | (~ spl80_310 | spl80_664)), inference(forward_demodulation, [], [f5073, f2992])).
fof(f5073, plain, (~ (h17(op1(e10, e13)) = op2(h17(e10), e20)) | spl80_664), inference(avatar_component_clause, [], [f5071])).
fof(f5071, plain, (spl80_664 <=> (h17(op1(e10, e13)) = op2(h17(e10), e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_664])])).
fof(f13560, plain, (~ spl80_76 | ~ spl80_132 | ~ spl80_300 | ~ spl80_310 | spl80_659), inference(avatar_contradiction_clause, [], [f13559])).
fof(f13559, plain, ($false | (~ spl80_76 | ~ spl80_132 | ~ spl80_300 | ~ spl80_310 | spl80_659)), inference(subsumption_resolution, [], [f13558, f1904])).
fof(f1904, plain, ((e21 = op2(e24, e23)) | ~ spl80_132), inference(avatar_component_clause, [], [f1902])).
fof(f1902, plain, (spl80_132 <=> (e21 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_132])])).
fof(f13558, plain, (~ (e21 = op2(e24, e23)) | (~ spl80_76 | ~ spl80_300 | ~ spl80_310 | spl80_659)), inference(forward_demodulation, [], [f13557, f2992])).
fof(f13557, plain, (~ (op2(e24, e23) = h17(e10)) | (~ spl80_76 | ~ spl80_300 | spl80_659)), inference(forward_demodulation, [], [f13556, f1619])).
fof(f1619, plain, ((e10 = op1(e11, e14)) | ~ spl80_76), inference(avatar_component_clause, [], [f1617])).
fof(f1617, plain, (spl80_76 <=> (e10 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_76])])).
fof(f13556, plain, (~ (op2(e24, e23) = h17(op1(e11, e14))) | (~ spl80_300 | spl80_659)), inference(forward_demodulation, [], [f5053, f2938])).
fof(f5053, plain, (~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | spl80_659), inference(avatar_component_clause, [], [f5051])).
fof(f5051, plain, (spl80_659 <=> (h17(op1(e11, e14)) = op2(e24, h17(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_659])])).
fof(f13405, plain, (~ spl80_34 | ~ spl80_246 | spl80_653), inference(avatar_contradiction_clause, [], [f13404])).
fof(f13404, plain, ($false | (~ spl80_34 | ~ spl80_246 | spl80_653)), inference(subsumption_resolution, [], [f13403, f2383])).
fof(f2383, plain, ((e20 = op2(e20, e20)) | ~ spl80_246), inference(avatar_component_clause, [], [f2381])).
fof(f2381, plain, (spl80_246 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_246])])).
fof(f13403, plain, (~ (e20 = op2(e20, e20)) | (~ spl80_34 | spl80_653)), inference(forward_demodulation, [], [f13402, f782])).
fof(f13402, plain, (~ (op2(e20, e20) = h17(e13)) | (~ spl80_34 | spl80_653)), inference(forward_demodulation, [], [f5029, f1442])).
fof(f1442, plain, ((e13 = op1(e13, e13)) | ~ spl80_34), inference(avatar_component_clause, [], [f1440])).
fof(f1440, plain, (spl80_34 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_34])])).
fof(f5029, plain, (~ (op2(e20, e20) = h17(op1(e13, e13))) | spl80_653), inference(avatar_component_clause, [], [f5027])).
fof(f5027, plain, (spl80_653 <=> (op2(e20, e20) = h17(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_653])])).
fof(f13242, plain, (~ spl80_7 | ~ spl80_175 | ~ spl80_300 | spl80_650), inference(avatar_contradiction_clause, [], [f13241])).
fof(f13241, plain, ($false | (~ spl80_7 | ~ spl80_175 | ~ spl80_300 | spl80_650)), inference(subsumption_resolution, [], [f13240, f2084])).
fof(f2084, plain, ((e24 = op2(e23, e20)) | ~ spl80_175), inference(avatar_component_clause, [], [f2082])).
fof(f2082, plain, (spl80_175 <=> (e24 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_175])])).
fof(f13240, plain, (~ (e24 = op2(e23, e20)) | (~ spl80_7 | ~ spl80_300 | spl80_650)), inference(forward_demodulation, [], [f13239, f781])).
fof(f13239, plain, (~ (op2(e23, e20) = h17(e11)) | (~ spl80_7 | ~ spl80_300 | spl80_650)), inference(forward_demodulation, [], [f13238, f1329])).
fof(f1329, plain, ((e11 = op1(e14, e13)) | ~ spl80_7), inference(avatar_component_clause, [], [f1327])).
fof(f1327, plain, (spl80_7 <=> (e11 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_7])])).
fof(f13238, plain, (~ (op2(e23, e20) = h17(op1(e14, e13))) | (~ spl80_300 | spl80_650)), inference(forward_demodulation, [], [f5017, f2938])).
fof(f5017, plain, (~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | spl80_650), inference(avatar_component_clause, [], [f5015])).
fof(f5015, plain, (spl80_650 <=> (h17(op1(e14, e13)) = op2(h17(e14), e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_650])])).
fof(f13069, plain, (~ spl80_24 | ~ spl80_166 | ~ spl80_300 | ~ spl80_310 | spl80_646), inference(avatar_contradiction_clause, [], [f13068])).
fof(f13068, plain, ($false | (~ spl80_24 | ~ spl80_166 | ~ spl80_300 | ~ spl80_310 | spl80_646)), inference(subsumption_resolution, [], [f13067, f2047])).
fof(f2047, plain, ((e20 = op2(e23, e21)) | ~ spl80_166), inference(avatar_component_clause, [], [f2045])).
fof(f2045, plain, (spl80_166 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_166])])).
fof(f13067, plain, (~ (e20 = op2(e23, e21)) | (~ spl80_24 | ~ spl80_300 | ~ spl80_310 | spl80_646)), inference(forward_demodulation, [], [f13066, f782])).
fof(f13066, plain, (~ (op2(e23, e21) = h17(e13)) | (~ spl80_24 | ~ spl80_300 | ~ spl80_310 | spl80_646)), inference(forward_demodulation, [], [f13065, f1400])).
fof(f1400, plain, ((e13 = op1(e14, e10)) | ~ spl80_24), inference(avatar_component_clause, [], [f1398])).
fof(f1398, plain, (spl80_24 <=> (e13 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_24])])).
fof(f13065, plain, (~ (op2(e23, e21) = h17(op1(e14, e10))) | (~ spl80_300 | ~ spl80_310 | spl80_646)), inference(forward_demodulation, [], [f13064, f2938])).
fof(f13064, plain, (~ (h17(op1(e14, e10)) = op2(h17(e14), e21)) | (~ spl80_310 | spl80_646)), inference(forward_demodulation, [], [f5001, f2992])).
fof(f5001, plain, (~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | spl80_646), inference(avatar_component_clause, [], [f4999])).
fof(f4999, plain, (spl80_646 <=> (h17(op1(e14, e10)) = op2(h17(e14), h17(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_646])])).
fof(f12892, plain, (~ spl80_54 | ~ spl80_181 | ~ spl80_300 | ~ spl80_305 | spl80_645), inference(avatar_contradiction_clause, [], [f12891])).
fof(f12891, plain, ($false | (~ spl80_54 | ~ spl80_181 | ~ spl80_300 | ~ spl80_305 | spl80_645)), inference(subsumption_resolution, [], [f12890, f2110])).
fof(f2110, plain, ((e20 = op2(e22, e23)) | ~ spl80_181), inference(avatar_component_clause, [], [f2108])).
fof(f2108, plain, (spl80_181 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_181])])).
fof(f12890, plain, (~ (e20 = op2(e22, e23)) | (~ spl80_54 | ~ spl80_300 | ~ spl80_305 | spl80_645)), inference(forward_demodulation, [], [f12889, f782])).
fof(f12889, plain, (~ (op2(e22, e23) = h17(e13)) | (~ spl80_54 | ~ spl80_300 | ~ spl80_305 | spl80_645)), inference(forward_demodulation, [], [f12888, f1526])).
fof(f1526, plain, ((e13 = op1(e12, e14)) | ~ spl80_54), inference(avatar_component_clause, [], [f1524])).
fof(f1524, plain, (spl80_54 <=> (e13 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_54])])).
fof(f12888, plain, (~ (op2(e22, e23) = h17(op1(e12, e14))) | (~ spl80_300 | ~ spl80_305 | spl80_645)), inference(forward_demodulation, [], [f12887, f2965])).
fof(f12887, plain, (~ (h17(op1(e12, e14)) = op2(h17(e12), e23)) | (~ spl80_300 | spl80_645)), inference(forward_demodulation, [], [f4997, f2938])).
fof(f4997, plain, (~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | spl80_645), inference(avatar_component_clause, [], [f4995])).
fof(f4995, plain, (spl80_645 <=> (h17(op1(e12, e14)) = op2(h17(e12), h17(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_645])])).
fof(f12717, plain, (~ spl80_63 | ~ spl80_188 | ~ spl80_305 | spl80_644), inference(avatar_contradiction_clause, [], [f12716])).
fof(f12716, plain, ($false | (~ spl80_63 | ~ spl80_188 | ~ spl80_305 | spl80_644)), inference(subsumption_resolution, [], [f12715, f2139])).
fof(f2139, plain, ((e22 = op2(e22, e22)) | ~ spl80_188), inference(avatar_component_clause, [], [f2137])).
fof(f2137, plain, (spl80_188 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_188])])).
fof(f12715, plain, (~ (e22 = op2(e22, e22)) | (~ spl80_63 | ~ spl80_305 | spl80_644)), inference(forward_demodulation, [], [f12714, f2965])).
fof(f12714, plain, (~ (op2(e22, e22) = h17(e12)) | (~ spl80_63 | ~ spl80_305 | spl80_644)), inference(forward_demodulation, [], [f12713, f1564])).
fof(f1564, plain, ((e12 = op1(e12, e12)) | ~ spl80_63), inference(avatar_component_clause, [], [f1562])).
fof(f1562, plain, (spl80_63 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_63])])).
fof(f12713, plain, (~ (op2(e22, e22) = h17(op1(e12, e12))) | (~ spl80_305 | spl80_644)), inference(forward_demodulation, [], [f4993, f2965])).
fof(f4993, plain, (~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | spl80_644), inference(avatar_component_clause, [], [f4991])).
fof(f4991, plain, (spl80_644 <=> (h17(op1(e12, e12)) = op2(h17(e12), h17(e12)))), introduced(avatar_definition, [new_symbols(naming, [spl80_644])])).
fof(f12545, plain, (~ spl80_92 | ~ spl80_130 | spl80_662), inference(avatar_contradiction_clause, [], [f12544])).
fof(f12544, plain, ($false | (~ spl80_92 | ~ spl80_130 | spl80_662)), inference(subsumption_resolution, [], [f12543, f1895])).
fof(f1895, plain, ((e24 = op2(e24, e24)) | ~ spl80_130), inference(avatar_component_clause, [], [f1893])).
fof(f1893, plain, (spl80_130 <=> (e24 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_130])])).
fof(f12543, plain, (~ (e24 = op2(e24, e24)) | (~ spl80_92 | spl80_662)), inference(forward_demodulation, [], [f12542, f781])).
fof(f12542, plain, (~ (op2(e24, e24) = h17(e11)) | (~ spl80_92 | spl80_662)), inference(forward_demodulation, [], [f5065, f1686])).
fof(f1686, plain, ((e11 = op1(e11, e11)) | ~ spl80_92), inference(avatar_component_clause, [], [f1684])).
fof(f1684, plain, (spl80_92 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_92])])).
fof(f5065, plain, (~ (op2(e24, e24) = h17(op1(e11, e11))) | spl80_662), inference(avatar_component_clause, [], [f5063])).
fof(f5063, plain, (spl80_662 <=> (op2(e24, e24) = h17(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_662])])).
fof(f12367, plain, (~ spl80_72 | ~ spl80_195 | ~ spl80_305 | ~ spl80_310 | spl80_643), inference(avatar_contradiction_clause, [], [f12366])).
fof(f12366, plain, ($false | (~ spl80_72 | ~ spl80_195 | ~ spl80_305 | ~ spl80_310 | spl80_643)), inference(subsumption_resolution, [], [f12365, f2168])).
fof(f2168, plain, ((e24 = op2(e22, e21)) | ~ spl80_195), inference(avatar_component_clause, [], [f2166])).
fof(f2166, plain, (spl80_195 <=> (e24 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_195])])).
fof(f12365, plain, (~ (e24 = op2(e22, e21)) | (~ spl80_72 | ~ spl80_305 | ~ spl80_310 | spl80_643)), inference(forward_demodulation, [], [f12364, f781])).
fof(f12364, plain, (~ (op2(e22, e21) = h17(e11)) | (~ spl80_72 | ~ spl80_305 | ~ spl80_310 | spl80_643)), inference(forward_demodulation, [], [f12363, f1602])).
fof(f1602, plain, ((e11 = op1(e12, e10)) | ~ spl80_72), inference(avatar_component_clause, [], [f1600])).
fof(f1600, plain, (spl80_72 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_72])])).
fof(f12363, plain, (~ (op2(e22, e21) = h17(op1(e12, e10))) | (~ spl80_305 | ~ spl80_310 | spl80_643)), inference(forward_demodulation, [], [f12362, f2965])).
fof(f12362, plain, (~ (h17(op1(e12, e10)) = op2(h17(e12), e21)) | (~ spl80_310 | spl80_643)), inference(forward_demodulation, [], [f4989, f2992])).
fof(f4989, plain, (~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | spl80_643), inference(avatar_component_clause, [], [f4987])).
fof(f4987, plain, (spl80_643 <=> (h17(op1(e12, e10)) = op2(h17(e12), h17(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_643])])).
fof(f12185, plain, (~ spl80_41 | spl80_655), inference(avatar_contradiction_clause, [], [f12184])).
fof(f12184, plain, ($false | (~ spl80_41 | spl80_655)), inference(trivial_inequality_removal, [], [f12183])).
fof(f12183, plain, (~ (h17(e10) = h17(e10)) | (~ spl80_41 | spl80_655)), inference(forward_demodulation, [], [f5037, f1472])).
fof(f1472, plain, ((e10 = op1(e13, e11)) | ~ spl80_41), inference(avatar_component_clause, [], [f1470])).
fof(f1470, plain, (spl80_41 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_41])])).
fof(f5037, plain, (~ (h17(e10) = h17(op1(e13, e11))) | spl80_655), inference(avatar_component_clause, [], [f5035])).
fof(f5035, plain, (spl80_655 <=> (h17(e10) = h17(op1(e13, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_655])])).
fof(f12034, plain, (spl80_300 | ~ spl80_310 | ~ spl80_445), inference(avatar_split_clause, [], [f12029, f3714, f2991, f2937])).
fof(f3714, plain, (spl80_445 <=> (e23 = h8(e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_445])])).
fof(f12029, plain, ((e23 = h17(e14)) | (~ spl80_310 | ~ spl80_445)), inference(backward_demodulation, [], [f11648, f3715])).
fof(f3715, plain, ((e23 = h8(e12)) | ~ spl80_445), inference(avatar_component_clause, [], [f3714])).
fof(f11648, plain, ((h8(e12) = h17(e14)) | ~ spl80_310), inference(forward_demodulation, [], [f11647, f2552])).
fof(f2552, plain, (h8(e12) = op2(e21, h8(e10))), inference(backward_demodulation, [], [f739, f738])).
fof(f738, plain, (op2(e24, e21) = h8(e10)), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ((h8(e14) = op2(op2(e24, e21), op2(e21, op2(e24, e21)))) & (h8(e12) = op2(e21, op2(e24, e21))) & (op2(e24, e21) = h8(e10)) & (e24 = h8(e13)) & (e21 = h8(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax23)).
fof(f739, plain, (h8(e12) = op2(e21, op2(e24, e21))), inference(cnf_transformation, [], [f23])).
fof(f11647, plain, ((h17(e14) = op2(e21, h8(e10))) | ~ spl80_310), inference(backward_demodulation, [], [f11282, f11646])).
fof(f11646, plain, ((h8(e10) = h17(e12)) | ~ spl80_310), inference(forward_demodulation, [], [f11645, f738])).
fof(f11645, plain, ((op2(e24, e21) = h17(e12)) | ~ spl80_310), inference(forward_demodulation, [], [f2646, f2992])).
fof(f2646, plain, (h17(e12) = op2(e24, h17(e10))), inference(backward_demodulation, [], [f784, f783])).
fof(f783, plain, (op2(e20, e24) = h17(e10)), inference(cnf_transformation, [], [f32])).
fof(f784, plain, (h17(e12) = op2(e24, op2(e20, e24))), inference(cnf_transformation, [], [f32])).
fof(f11282, plain, ((h17(e14) = op2(e21, h17(e12))) | ~ spl80_310), inference(backward_demodulation, [], [f2649, f2992])).
fof(f2649, plain, (h17(e14) = op2(h17(e10), h17(e12))), inference(backward_demodulation, [], [f2639, f783])).
fof(f2639, plain, (h17(e14) = op2(op2(e20, e24), h17(e12))), inference(backward_demodulation, [], [f785, f784])).
fof(f785, plain, (h17(e14) = op2(op2(e20, e24), op2(e24, op2(e20, e24)))), inference(cnf_transformation, [], [f32])).
fof(f12031, plain, (~ spl80_102 | ~ spl80_210 | ~ spl80_310 | ~ spl80_445 | spl80_642), inference(avatar_contradiction_clause, [], [f12030])).
fof(f12030, plain, ($false | (~ spl80_102 | ~ spl80_210 | ~ spl80_310 | ~ spl80_445 | spl80_642)), inference(subsumption_resolution, [], [f12027, f2231])).
fof(f2231, plain, ((e24 = op2(e21, e23)) | ~ spl80_210), inference(avatar_component_clause, [], [f2229])).
fof(f2229, plain, (spl80_210 <=> (e24 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_210])])).
fof(f12027, plain, (~ (e24 = op2(e21, e23)) | (~ spl80_102 | ~ spl80_310 | ~ spl80_445 | spl80_642)), inference(backward_demodulation, [], [f12002, f3715])).
fof(f12002, plain, (~ (e24 = op2(e21, h8(e12))) | (~ spl80_102 | ~ spl80_310 | spl80_642)), inference(forward_demodulation, [], [f12001, f781])).
fof(f12001, plain, (~ (h17(e11) = op2(e21, h8(e12))) | (~ spl80_102 | ~ spl80_310 | spl80_642)), inference(forward_demodulation, [], [f12000, f1728])).
fof(f1728, plain, ((e11 = op1(e10, e14)) | ~ spl80_102), inference(avatar_component_clause, [], [f1726])).
fof(f1726, plain, (spl80_102 <=> (e11 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_102])])).
fof(f12000, plain, (~ (h17(op1(e10, e14)) = op2(e21, h8(e12))) | (~ spl80_310 | spl80_642)), inference(forward_demodulation, [], [f11999, f2992])).
fof(f11999, plain, (~ (h17(op1(e10, e14)) = op2(h17(e10), h8(e12))) | (~ spl80_310 | spl80_642)), inference(forward_demodulation, [], [f4985, f11648])).
fof(f4985, plain, (~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | spl80_642), inference(avatar_component_clause, [], [f4983])).
fof(f4983, plain, (spl80_642 <=> (h17(op1(e10, e14)) = op2(h17(e10), h17(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_642])])).
fof(f11930, plain, (spl80_179 | ~ spl80_370 | ~ spl80_1058), inference(avatar_split_clause, [], [f11663, f7902, f3313, f2099])).
fof(f3313, plain, (spl80_370 <=> (e22 = h13(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_370])])).
fof(f7902, plain, (spl80_1058 <=> (e24 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_1058])])).
fof(f11663, plain, ((e23 = op2(e22, e24)) | (~ spl80_370 | ~ spl80_1058)), inference(backward_demodulation, [], [f11250, f7903])).
fof(f7903, plain, ((e24 = h3(e10)) | ~ spl80_1058), inference(avatar_component_clause, [], [f7902])).
fof(f11250, plain, ((e23 = op2(e22, h3(e10))) | ~ spl80_370), inference(backward_demodulation, [], [f2608, f3314])).
fof(f3314, plain, ((e22 = h13(e10)) | ~ spl80_370), inference(avatar_component_clause, [], [f3313])).
fof(f2608, plain, (e23 = op2(h13(e10), h3(e10))), inference(backward_demodulation, [], [f2484, f763])).
fof(f763, plain, (op2(e20, e23) = h13(e10)), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ((h13(e14) = op2(op2(e20, e23), op2(e23, op2(e20, e23)))) & (h13(e12) = op2(e23, op2(e20, e23))) & (op2(e20, e23) = h13(e10)) & (e20 = h13(e13)) & (e23 = h13(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax28)).
fof(f2484, plain, (e23 = op2(op2(e20, e23), h3(e10))), inference(backward_demodulation, [], [f390, f713])).
fof(f713, plain, (op2(e23, e20) = h3(e10)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ((h3(e14) = op2(op2(e23, e20), op2(e20, op2(e23, e20)))) & (h3(e12) = op2(e20, op2(e23, e20))) & (op2(e23, e20) = h3(e10)) & (e23 = h3(e13)) & (e20 = h3(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax18)).
fof(f390, plain, (e23 = op2(op2(e20, e23), op2(e23, e20))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ((e24 = op2(op2(e24, e24), op2(e24, e24))) & (e24 = op2(op2(e23, e24), op2(e24, e23))) & (e24 = op2(op2(e22, e24), op2(e24, e22))) & (e24 = op2(op2(e21, e24), op2(e24, e21))) & (e24 = op2(op2(e20, e24), op2(e24, e20))) & (e23 = op2(op2(e24, e23), op2(e23, e24))) & (e23 = op2(op2(e23, e23), op2(e23, e23))) & (e23 = op2(op2(e22, e23), op2(e23, e22))) & (e23 = op2(op2(e21, e23), op2(e23, e21))) & (e23 = op2(op2(e20, e23), op2(e23, e20))) & (e22 = op2(op2(e24, e22), op2(e22, e24))) & (e22 = op2(op2(e23, e22), op2(e22, e23))) & (e22 = op2(op2(e22, e22), op2(e22, e22))) & (e22 = op2(op2(e21, e22), op2(e22, e21))) & (e22 = op2(op2(e20, e22), op2(e22, e20))) & (e21 = op2(op2(e24, e21), op2(e21, e24))) & (e21 = op2(op2(e23, e21), op2(e21, e23))) & (e21 = op2(op2(e22, e21), op2(e21, e22))) & (e21 = op2(op2(e21, e21), op2(e21, e21))) & (e21 = op2(op2(e20, e21), op2(e21, e20))) & (e20 = op2(op2(e24, e20), op2(e20, e24))) & (e20 = op2(op2(e23, e20), op2(e20, e23))) & (e20 = op2(op2(e22, e20), op2(e20, e22))) & (e20 = op2(op2(e21, e20), op2(e20, e21))) & (e20 = op2(op2(e20, e20), op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax6)).
fof(f11913, plain, (~ spl80_7 | ~ spl80_9), inference(avatar_contradiction_clause, [], [f11912])).
fof(f11912, plain, ($false | (~ spl80_7 | ~ spl80_9)), inference(subsumption_resolution, [], [f11911, f605])).
fof(f605, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (~ (e13 = e14) & ~ (e12 = e14) & ~ (e12 = e13) & ~ (e11 = e14) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e14) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax9)).
fof(f11911, plain, ((e11 = e13) | (~ spl80_7 | ~ spl80_9)), inference(forward_demodulation, [], [f1337, f1329])).
fof(f1337, plain, ((e13 = op1(e14, e13)) | ~ spl80_9), inference(avatar_component_clause, [], [f1335])).
fof(f1335, plain, (spl80_9 <=> (e13 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_9])])).
fof(f11905, plain, (~ spl80_76 | ~ spl80_79), inference(avatar_contradiction_clause, [], [f11904])).
fof(f11904, plain, ($false | (~ spl80_76 | ~ spl80_79)), inference(subsumption_resolution, [], [f11903, f602])).
fof(f602, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f9])).
fof(f11903, plain, ((e10 = e13) | (~ spl80_76 | ~ spl80_79)), inference(forward_demodulation, [], [f1631, f1619])).
fof(f1631, plain, ((e13 = op1(e11, e14)) | ~ spl80_79), inference(avatar_component_clause, [], [f1629])).
fof(f1629, plain, (spl80_79 <=> (e13 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_79])])).
fof(f11878, plain, (~ spl80_158 | ~ spl80_181), inference(avatar_contradiction_clause, [], [f11877])).
fof(f11877, plain, ($false | (~ spl80_158 | ~ spl80_181)), inference(subsumption_resolution, [], [f11876, f612])).
fof(f612, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f10])).
fof(f10, plain, (~ (e23 = e24) & ~ (e22 = e24) & ~ (e22 = e23) & ~ (e21 = e24) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e24) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax10)).
fof(f11876, plain, ((e20 = e23) | (~ spl80_158 | ~ spl80_181)), inference(forward_demodulation, [], [f11872, f2110])).
fof(f11872, plain, ((e23 = op2(e22, e23)) | ~ spl80_158), inference(backward_demodulation, [], [f688, f2013])).
fof(f2013, plain, ((e22 = op2(e23, e23)) | ~ spl80_158), inference(avatar_component_clause, [], [f2011])).
fof(f2011, plain, (spl80_158 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_158])])).
fof(f688, plain, (e23 = op2(op2(e23, e23), e23)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e24 = op2(op2(e24, e24), e24)) & (e24 = op2(op2(e24, e23), e23)) & (e24 = op2(op2(e24, e22), e22)) & (e24 = op2(op2(e24, e21), e21)) & (e24 = op2(op2(e24, e20), e20)) & (e23 = op2(op2(e23, e24), e24)) & (e23 = op2(op2(e23, e23), e23)) & (e23 = op2(op2(e23, e22), e22)) & (e23 = op2(op2(e23, e21), e21)) & (e23 = op2(op2(e23, e20), e20)) & (e22 = op2(op2(e22, e24), e24)) & (e22 = op2(op2(e22, e23), e23)) & (e22 = op2(op2(e22, e22), e22)) & (e22 = op2(op2(e22, e21), e21)) & (e22 = op2(op2(e22, e20), e20)) & (e21 = op2(op2(e21, e24), e24)) & (e21 = op2(op2(e21, e23), e23)) & (e21 = op2(op2(e21, e22), e22)) & (e21 = op2(op2(e21, e21), e21)) & (e21 = op2(op2(e21, e20), e20)) & (e20 = op2(op2(e20, e24), e24)) & (e20 = op2(op2(e20, e23), e23)) & (e20 = op2(op2(e20, e22), e22)) & (e20 = op2(op2(e20, e21), e21)) & (e20 = op2(op2(e20, e20), e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax13)).
fof(f11858, plain, (~ spl80_209 | ~ spl80_210), inference(avatar_contradiction_clause, [], [f11857])).
fof(f11857, plain, ($false | (~ spl80_209 | ~ spl80_210)), inference(subsumption_resolution, [], [f11856, f619])).
fof(f619, plain, ~ (e23 = e24), inference(cnf_transformation, [], [f10])).
fof(f11856, plain, ((e23 = e24) | (~ spl80_209 | ~ spl80_210)), inference(backward_demodulation, [], [f2231, f2227])).
fof(f2227, plain, ((e23 = op2(e21, e23)) | ~ spl80_209), inference(avatar_component_clause, [], [f2225])).
fof(f2225, plain, (spl80_209 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_209])])).
fof(f11816, plain, (spl80_305 | ~ spl80_310 | ~ spl80_450), inference(avatar_split_clause, [], [f11810, f3741, f2991, f2964])).
fof(f3741, plain, (spl80_450 <=> (e22 = h8(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_450])])).
fof(f11810, plain, ((e22 = h17(e12)) | (~ spl80_310 | ~ spl80_450)), inference(backward_demodulation, [], [f11646, f3742])).
fof(f3742, plain, ((e22 = h8(e10)) | ~ spl80_450), inference(avatar_component_clause, [], [f3741])).
fof(f11815, plain, (spl80_445 | ~ spl80_214 | ~ spl80_450), inference(avatar_split_clause, [], [f11814, f3741, f2246, f3714])).
fof(f2246, plain, (spl80_214 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_214])])).
fof(f11814, plain, ((e23 = h8(e12)) | (~ spl80_214 | ~ spl80_450)), inference(forward_demodulation, [], [f11804, f2248])).
fof(f2248, plain, ((e23 = op2(e21, e22)) | ~ spl80_214), inference(avatar_component_clause, [], [f2246])).
fof(f11804, plain, ((op2(e21, e22) = h8(e12)) | ~ spl80_450), inference(backward_demodulation, [], [f2552, f3742])).
fof(f11813, plain, (spl80_195 | ~ spl80_450), inference(avatar_split_clause, [], [f11803, f3741, f2166])).
fof(f11803, plain, ((e24 = op2(e22, e21)) | ~ spl80_450), inference(backward_demodulation, [], [f2551, f3742])).
fof(f2551, plain, (e24 = op2(h8(e10), e21)), inference(backward_demodulation, [], [f691, f738])).
fof(f691, plain, (e24 = op2(op2(e24, e21), e21)), inference(cnf_transformation, [], [f13])).
fof(f11798, plain, (spl80_210 | ~ spl80_310 | ~ spl80_496), inference(avatar_split_clause, [], [f11788, f3999, f2991, f2229])).
fof(f11788, plain, ((e24 = op2(e21, e23)) | (~ spl80_310 | ~ spl80_496)), inference(backward_demodulation, [], [f11649, f4000])).
fof(f11649, plain, ((e24 = op2(e21, h4(e10))) | ~ spl80_310), inference(forward_demodulation, [], [f2648, f2992])).
fof(f2648, plain, (e24 = op2(h17(e10), h4(e10))), inference(backward_demodulation, [], [f2496, f783])).
fof(f2496, plain, (e24 = op2(op2(e20, e24), h4(e10))), inference(backward_demodulation, [], [f395, f718])).
fof(f718, plain, (op2(e24, e20) = h4(e10)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ((h4(e14) = op2(op2(e24, e20), op2(e20, op2(e24, e20)))) & (h4(e12) = op2(e20, op2(e24, e20))) & (op2(e24, e20) = h4(e10)) & (e24 = h4(e13)) & (e20 = h4(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax19)).
fof(f395, plain, (e24 = op2(op2(e20, e24), op2(e24, e20))), inference(cnf_transformation, [], [f6])).
fof(f11792, plain, (~ spl80_134 | ~ spl80_496), inference(avatar_split_clause, [], [f11780, f3999, f1910])).
fof(f1910, plain, (spl80_134 <=> (e23 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_134])])).
fof(f11780, plain, (~ (e23 = op2(e24, e23)) | ~ spl80_496), inference(backward_demodulation, [], [f2500, f4000])).
fof(f2500, plain, ~ (op2(e24, e23) = h4(e10)), inference(backward_demodulation, [], [f592, f718])).
fof(f592, plain, ~ (op2(e24, e20) = op2(e24, e23)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (op2(e24, e23) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e24)) & ~ (op2(e24, e21) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e24)) & ~ (op2(e24, e20) = op2(e24, e23)) & ~ (op2(e24, e20) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e21)) & ~ (op2(e23, e23) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e24)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e24)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e23) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e24)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e24)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e23) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e24)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e24)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e23) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e24)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e24)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e23, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e24, e24)) & ~ (op2(e21, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e24, e24)) & ~ (op2(e20, e24) = op2(e23, e24)) & ~ (op2(e20, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e21, e24)) & ~ (op2(e23, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e24, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e24, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e23, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e24, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e24, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e23, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e24, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e24, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e23, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e24, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e24, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax8)).
fof(f11761, plain, (~ spl80_197 | ~ spl80_394 | ~ spl80_546 | ~ spl80_548 | ~ spl80_892 | ~ spl80_1113), inference(avatar_contradiction_clause, [], [f11760])).
fof(f11760, plain, ($false | (~ spl80_197 | ~ spl80_394 | ~ spl80_546 | ~ spl80_548 | ~ spl80_892 | ~ spl80_1113)), inference(subsumption_resolution, [], [f11759, f616])).
fof(f616, plain, ~ (e21 = e24), inference(cnf_transformation, [], [f10])).
fof(f11759, plain, ((e21 = e24) | (~ spl80_197 | ~ spl80_394 | ~ spl80_546 | ~ spl80_548 | ~ spl80_892 | ~ spl80_1113)), inference(forward_demodulation, [], [f11749, f2177])).
fof(f2177, plain, ((e21 = op2(e22, e20)) | ~ spl80_197), inference(avatar_component_clause, [], [f2175])).
fof(f2175, plain, (spl80_197 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_197])])).
fof(f11749, plain, ((e24 = op2(e22, e20)) | (~ spl80_394 | ~ spl80_546 | ~ spl80_548 | ~ spl80_892 | ~ spl80_1113)), inference(backward_demodulation, [], [f11242, f4266])).
fof(f4266, plain, ((e22 = h1(e14)) | ~ spl80_546), inference(avatar_component_clause, [], [f4265])).
fof(f4265, plain, (spl80_546 <=> (e22 = h1(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_546])])).
fof(f11242, plain, ((e24 = op2(h1(e14), e20)) | (~ spl80_394 | ~ spl80_548 | ~ spl80_892 | ~ spl80_1113)), inference(backward_demodulation, [], [f11111, f3442])).
fof(f3442, plain, ((e20 = h12(e10)) | ~ spl80_394), inference(avatar_component_clause, [], [f3441])).
fof(f3441, plain, (spl80_394 <=> (e20 = h12(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_394])])).
fof(f11111, plain, ((e24 = op2(h1(e14), h12(e10))) | (~ spl80_548 | ~ spl80_892 | ~ spl80_1113)), inference(backward_demodulation, [], [f11025, f11105])).
fof(f11105, plain, ((h1(e14) = h9(e12)) | (~ spl80_548 | ~ spl80_892 | ~ spl80_1113)), inference(backward_demodulation, [], [f10508, f11102])).
fof(f11102, plain, ((op2(e22, e24) = h1(e14)) | (~ spl80_548 | ~ spl80_1113)), inference(backward_demodulation, [], [f10526, f8246])).
fof(f8246, plain, ((e24 = h1(e12)) | ~ spl80_1113), inference(avatar_component_clause, [], [f8245])).
fof(f8245, plain, (spl80_1113 <=> (e24 = h1(e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_1113])])).
fof(f10526, plain, ((h1(e14) = op2(e22, h1(e12))) | ~ spl80_548), inference(backward_demodulation, [], [f2468, f4278])).
fof(f4278, plain, ((e22 = h1(e10)) | ~ spl80_548), inference(avatar_component_clause, [], [f4277])).
fof(f4277, plain, (spl80_548 <=> (e22 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_548])])).
fof(f2468, plain, (h1(e14) = op2(h1(e10), h1(e12))), inference(backward_demodulation, [], [f2455, f703])).
fof(f703, plain, (op2(e21, e20) = h1(e10)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((h1(e14) = op2(op2(e21, e20), op2(e20, op2(e21, e20)))) & (h1(e12) = op2(e20, op2(e21, e20))) & (op2(e21, e20) = h1(e10)) & (e21 = h1(e13)) & (e20 = h1(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax16)).
fof(f2455, plain, (h1(e14) = op2(op2(e21, e20), h1(e12))), inference(backward_demodulation, [], [f705, f704])).
fof(f704, plain, (h1(e12) = op2(e20, op2(e21, e20))), inference(cnf_transformation, [], [f16])).
fof(f705, plain, (h1(e14) = op2(op2(e21, e20), op2(e20, op2(e21, e20)))), inference(cnf_transformation, [], [f16])).
fof(f10508, plain, ((op2(e22, e24) = h9(e12)) | ~ spl80_892), inference(backward_demodulation, [], [f2563, f6698])).
fof(f6698, plain, ((e24 = h9(e10)) | ~ spl80_892), inference(avatar_component_clause, [], [f6697])).
fof(f6697, plain, (spl80_892 <=> (e24 = h9(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_892])])).
fof(f2563, plain, (h9(e12) = op2(e22, h9(e10))), inference(backward_demodulation, [], [f744, f743])).
fof(f743, plain, (op2(e20, e22) = h9(e10)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ((h9(e14) = op2(op2(e20, e22), op2(e22, op2(e20, e22)))) & (h9(e12) = op2(e22, op2(e20, e22))) & (op2(e20, e22) = h9(e10)) & (e20 = h9(e13)) & (e22 = h9(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax24)).
fof(f744, plain, (h9(e12) = op2(e22, op2(e20, e22))), inference(cnf_transformation, [], [f24])).
fof(f11025, plain, ((e24 = op2(h9(e12), h12(e10))) | ~ spl80_892), inference(backward_demodulation, [], [f2667, f11020])).
fof(f11020, plain, ((h9(e12) = h19(e10)) | ~ spl80_892), inference(forward_demodulation, [], [f793, f10508])).
fof(f793, plain, (op2(e22, e24) = h19(e10)), inference(cnf_transformation, [], [f34])).
fof(f34, plain, ((h19(e14) = op2(op2(e22, e24), op2(e24, op2(e22, e24)))) & (h19(e12) = op2(e24, op2(e22, e24))) & (op2(e22, e24) = h19(e10)) & (e22 = h19(e13)) & (e24 = h19(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax34)).
fof(f2667, plain, (e24 = op2(h19(e10), h12(e10))), inference(backward_demodulation, [], [f2591, f793])).
fof(f2591, plain, (e24 = op2(op2(e22, e24), h12(e10))), inference(backward_demodulation, [], [f397, f758])).
fof(f758, plain, (op2(e24, e22) = h12(e10)), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ((h12(e14) = op2(op2(e24, e22), op2(e22, op2(e24, e22)))) & (h12(e12) = op2(e22, op2(e24, e22))) & (op2(e24, e22) = h12(e10)) & (e24 = h12(e13)) & (e22 = h12(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax27)).
fof(f397, plain, (e24 = op2(op2(e22, e24), op2(e24, e22))), inference(cnf_transformation, [], [f6])).
fof(f11729, plain, (~ spl80_56 | ~ spl80_298 | ~ spl80_310 | spl80_657), inference(avatar_contradiction_clause, [], [f11728])).
fof(f11728, plain, ($false | (~ spl80_56 | ~ spl80_298 | ~ spl80_310 | spl80_657)), inference(subsumption_resolution, [], [f11727, f2992])).
fof(f11727, plain, (~ (e21 = h17(e10)) | (~ spl80_56 | ~ spl80_298 | ~ spl80_310 | spl80_657)), inference(forward_demodulation, [], [f11726, f1535])).
fof(f1535, plain, ((e10 = op1(e12, e13)) | ~ spl80_56), inference(avatar_component_clause, [], [f1533])).
fof(f1533, plain, (spl80_56 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_56])])).
fof(f11726, plain, (~ (e21 = h17(op1(e12, e13))) | (~ spl80_298 | ~ spl80_310 | spl80_657)), inference(forward_demodulation, [], [f11725, f11624])).
fof(f11624, plain, ((e21 = op2(h8(e10), e20)) | ~ spl80_298), inference(forward_demodulation, [], [f2657, f2929])).
fof(f2929, plain, ((e20 = h18(e10)) | ~ spl80_298), inference(avatar_component_clause, [], [f2928])).
fof(f2928, plain, (spl80_298 <=> (e20 = h18(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_298])])).
fof(f2657, plain, (e21 = op2(h8(e10), h18(e10))), inference(backward_demodulation, [], [f2545, f788])).
fof(f788, plain, (op2(e21, e24) = h18(e10)), inference(cnf_transformation, [], [f33])).
fof(f33, plain, ((h18(e14) = op2(op2(e21, e24), op2(e24, op2(e21, e24)))) & (h18(e12) = op2(e24, op2(e21, e24))) & (op2(e21, e24) = h18(e10)) & (e21 = h18(e13)) & (e24 = h18(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax33)).
fof(f2545, plain, (e21 = op2(h8(e10), op2(e21, e24))), inference(backward_demodulation, [], [f384, f738])).
fof(f384, plain, (e21 = op2(op2(e24, e21), op2(e21, e24))), inference(cnf_transformation, [], [f6])).
fof(f11725, plain, (~ (h17(op1(e12, e13)) = op2(h8(e10), e20)) | (~ spl80_310 | spl80_657)), inference(forward_demodulation, [], [f5045, f11646])).
fof(f5045, plain, (~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | spl80_657), inference(avatar_component_clause, [], [f5043])).
fof(f5043, plain, (spl80_657 <=> (h17(op1(e12, e13)) = op2(h17(e12), e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_657])])).
fof(f11669, plain, (spl80_149 | ~ spl80_1058), inference(avatar_split_clause, [], [f11658, f7902, f1973])).
fof(f1973, plain, (spl80_149 <=> (e23 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_149])])).
fof(f11658, plain, ((e23 = op2(e24, e20)) | ~ spl80_1058), inference(backward_demodulation, [], [f2491, f7903])).
fof(f2491, plain, (e23 = op2(h3(e10), e20)), inference(backward_demodulation, [], [f685, f713])).
fof(f685, plain, (e23 = op2(op2(e23, e20), e20)), inference(cnf_transformation, [], [f13])).
fof(f11580, plain, (~ spl80_147 | ~ spl80_536), inference(avatar_split_clause, [], [f11579, f4213, f1965])).
fof(f1965, plain, (spl80_147 <=> (e21 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_147])])).
fof(f4213, plain, (spl80_536 <=> (e21 = h2(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_536])])).
fof(f11579, plain, (~ (e21 = op2(e24, e20)) | ~ spl80_536), inference(forward_demodulation, [], [f2474, f4214])).
fof(f4214, plain, ((e21 = h2(e10)) | ~ spl80_536), inference(avatar_component_clause, [], [f4213])).
fof(f2474, plain, ~ (op2(e24, e20) = h2(e10)), inference(backward_demodulation, [], [f508, f708])).
fof(f708, plain, (op2(e22, e20) = h2(e10)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((h2(e14) = op2(op2(e22, e20), op2(e20, op2(e22, e20)))) & (h2(e12) = op2(e20, op2(e22, e20))) & (op2(e22, e20) = h2(e10)) & (e22 = h2(e13)) & (e20 = h2(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax17)).
fof(f508, plain, ~ (op2(e22, e20) = op2(e24, e20)), inference(cnf_transformation, [], [f8])).
fof(f11577, plain, (spl80_143 | ~ spl80_536 | ~ spl80_892), inference(avatar_split_clause, [], [f11165, f6697, f4213, f1948])).
fof(f1948, plain, (spl80_143 <=> (e22 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_143])])).
fof(f11165, plain, ((e22 = op2(e24, e21)) | (~ spl80_536 | ~ spl80_892)), inference(backward_demodulation, [], [f10510, f4214])).
fof(f10510, plain, ((e22 = op2(e24, h2(e10))) | ~ spl80_892), inference(backward_demodulation, [], [f2565, f6698])).
fof(f2565, plain, (e22 = op2(h9(e10), h2(e10))), inference(backward_demodulation, [], [f2471, f743])).
fof(f2471, plain, (e22 = op2(op2(e20, e22), h2(e10))), inference(backward_demodulation, [], [f385, f708])).
fof(f385, plain, (e22 = op2(op2(e20, e22), op2(e22, e20))), inference(cnf_transformation, [], [f6])).
fof(f11551, plain, (~ spl80_87 | ~ spl80_37), inference(avatar_split_clause, [], [f11550, f1453, f1663])).
fof(f1663, plain, (spl80_87 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_87])])).
fof(f11550, plain, (~ (e11 = op1(e11, e12)) | ~ spl80_37), inference(forward_demodulation, [], [f425, f1455])).
fof(f425, plain, ~ (op1(e11, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (op1(e14, e13) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e14)) & ~ (op1(e14, e11) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e14)) & ~ (op1(e14, e10) = op1(e14, e13)) & ~ (op1(e14, e10) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e11)) & ~ (op1(e13, e13) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e14)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e14)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e13) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e14)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e14)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e13) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e14)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e14)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e13) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e14)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e14)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e13, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e14, e14)) & ~ (op1(e11, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e14, e14)) & ~ (op1(e10, e14) = op1(e13, e14)) & ~ (op1(e10, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e11, e14)) & ~ (op1(e13, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e14, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e14, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e13, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e14, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e14, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e13, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e14, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e14, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e13, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e14, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e14, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax7)).
fof(f11542, plain, (spl80_89 | ~ spl80_37), inference(avatar_split_clause, [], [f10805, f1453, f1671])).
fof(f10805, plain, ((e13 = op1(e11, e12)) | ~ spl80_37), inference(backward_demodulation, [], [f662, f1455])).
fof(f662, plain, (e13 = op1(op1(e13, e12), e12)), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e14 = op1(op1(e14, e14), e14)) & (e14 = op1(op1(e14, e13), e13)) & (e14 = op1(op1(e14, e12), e12)) & (e14 = op1(op1(e14, e11), e11)) & (e14 = op1(op1(e14, e10), e10)) & (e13 = op1(op1(e13, e14), e14)) & (e13 = op1(op1(e13, e13), e13)) & (e13 = op1(op1(e13, e12), e12)) & (e13 = op1(op1(e13, e11), e11)) & (e13 = op1(op1(e13, e10), e10)) & (e12 = op1(op1(e12, e14), e14)) & (e12 = op1(op1(e12, e13), e13)) & (e12 = op1(op1(e12, e12), e12)) & (e12 = op1(op1(e12, e11), e11)) & (e12 = op1(op1(e12, e10), e10)) & (e11 = op1(op1(e11, e14), e14)) & (e11 = op1(op1(e11, e13), e13)) & (e11 = op1(op1(e11, e12), e12)) & (e11 = op1(op1(e11, e11), e11)) & (e11 = op1(op1(e11, e10), e10)) & (e10 = op1(op1(e10, e14), e14)) & (e10 = op1(op1(e10, e13), e13)) & (e10 = op1(op1(e10, e12), e12)) & (e10 = op1(op1(e10, e11), e11)) & (e10 = op1(op1(e10, e10), e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax12)).
fof(f11509, plain, (~ spl80_7 | ~ spl80_35), inference(avatar_contradiction_clause, [], [f11508])).
fof(f11508, plain, ($false | (~ spl80_7 | ~ spl80_35)), inference(subsumption_resolution, [], [f11507, f605])).
fof(f11507, plain, ((e11 = e13) | (~ spl80_7 | ~ spl80_35)), inference(forward_demodulation, [], [f11506, f1329])).
fof(f11506, plain, ((e13 = op1(e14, e13)) | ~ spl80_35), inference(forward_demodulation, [], [f663, f1446])).
fof(f1446, plain, ((e14 = op1(e13, e13)) | ~ spl80_35), inference(avatar_component_clause, [], [f1444])).
fof(f1444, plain, (spl80_35 <=> (e14 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_35])])).
fof(f663, plain, (e13 = op1(op1(e13, e13), e13)), inference(cnf_transformation, [], [f12])).
fof(f11489, plain, (~ spl80_4 | ~ spl80_24), inference(avatar_split_clause, [], [f11484, f1398, f1314])).
fof(f1314, plain, (spl80_4 <=> (e13 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_4])])).
fof(f11484, plain, (~ (e13 = op1(e14, e14)) | ~ spl80_24), inference(backward_demodulation, [], [f493, f1400])).
fof(f493, plain, ~ (op1(e14, e10) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f11482, plain, (~ spl80_3 | ~ spl80_28), inference(avatar_split_clause, [], [f11476, f1415, f1310])).
fof(f1310, plain, (spl80_3 <=> (e12 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_3])])).
fof(f11476, plain, (~ (e12 = op1(e14, e14)) | ~ spl80_28), inference(backward_demodulation, [], [f449, f1417])).
fof(f449, plain, ~ (op1(e13, e14) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f11468, plain, (spl80_24 | ~ spl80_50), inference(avatar_split_clause, [], [f11463, f1507, f1398])).
fof(f11463, plain, ((e13 = op1(e14, e10)) | ~ spl80_50), inference(backward_demodulation, [], [f660, f1509])).
fof(f660, plain, (e13 = op1(op1(e13, e10), e10)), inference(cnf_transformation, [], [f12])).
fof(f11453, plain, (spl80_28 | ~ spl80_70 | ~ spl80_89), inference(avatar_split_clause, [], [f11449, f1671, f1591, f1415])).
fof(f11449, plain, ((e12 = op1(e13, e14)) | (~ spl80_70 | ~ spl80_89)), inference(backward_demodulation, [], [f10374, f1593])).
fof(f10374, plain, ((e12 = op1(e13, op1(e12, e11))) | ~ spl80_89), inference(forward_demodulation, [], [f286, f1673])).
fof(f286, plain, (e12 = op1(op1(e11, e12), op1(e12, e11))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, ((e14 = op1(op1(e14, e14), op1(e14, e14))) & (e14 = op1(op1(e13, e14), op1(e14, e13))) & (e14 = op1(op1(e12, e14), op1(e14, e12))) & (e14 = op1(op1(e11, e14), op1(e14, e11))) & (e14 = op1(op1(e10, e14), op1(e14, e10))) & (e13 = op1(op1(e14, e13), op1(e13, e14))) & (e13 = op1(op1(e13, e13), op1(e13, e13))) & (e13 = op1(op1(e12, e13), op1(e13, e12))) & (e13 = op1(op1(e11, e13), op1(e13, e11))) & (e13 = op1(op1(e10, e13), op1(e13, e10))) & (e12 = op1(op1(e14, e12), op1(e12, e14))) & (e12 = op1(op1(e13, e12), op1(e12, e13))) & (e12 = op1(op1(e12, e12), op1(e12, e12))) & (e12 = op1(op1(e11, e12), op1(e12, e11))) & (e12 = op1(op1(e10, e12), op1(e12, e10))) & (e11 = op1(op1(e14, e11), op1(e11, e14))) & (e11 = op1(op1(e13, e11), op1(e11, e13))) & (e11 = op1(op1(e12, e11), op1(e11, e12))) & (e11 = op1(op1(e11, e11), op1(e11, e11))) & (e11 = op1(op1(e10, e11), op1(e11, e10))) & (e10 = op1(op1(e14, e10), op1(e10, e14))) & (e10 = op1(op1(e13, e10), op1(e10, e13))) & (e10 = op1(op1(e12, e10), op1(e10, e12))) & (e10 = op1(op1(e11, e10), op1(e10, e11))) & (e10 = op1(op1(e10, e10), op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax3)).
fof(f11452, plain, (spl80_18 | ~ spl80_70), inference(avatar_split_clause, [], [f11448, f1591, f1373])).
fof(f11448, plain, ((e12 = op1(e14, e11)) | ~ spl80_70), inference(backward_demodulation, [], [f656, f1593])).
fof(f656, plain, (e12 = op1(op1(e12, e11), e11)), inference(cnf_transformation, [], [f12])).
fof(f11438, plain, (~ spl80_41 | ~ spl80_82 | ~ spl80_98), inference(avatar_contradiction_clause, [], [f11437])).
fof(f11437, plain, ($false | (~ spl80_41 | ~ spl80_82 | ~ spl80_98)), inference(subsumption_resolution, [], [f11436, f607])).
fof(f607, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f9])).
fof(f11436, plain, ((e12 = e13) | (~ spl80_41 | ~ spl80_82 | ~ spl80_98)), inference(forward_demodulation, [], [f11434, f1711])).
fof(f11434, plain, ((e13 = op1(e11, e10)) | (~ spl80_41 | ~ spl80_82)), inference(backward_demodulation, [], [f9199, f1644])).
fof(f1644, plain, ((e11 = op1(e11, e13)) | ~ spl80_82), inference(avatar_component_clause, [], [f1642])).
fof(f1642, plain, (spl80_82 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_82])])).
fof(f9199, plain, ((e13 = op1(op1(e11, e13), e10)) | ~ spl80_41), inference(forward_demodulation, [], [f291, f1472])).
fof(f291, plain, (e13 = op1(op1(e11, e13), op1(e13, e11))), inference(cnf_transformation, [], [f3])).
fof(f11364, plain, (spl80_546 | ~ spl80_178 | ~ spl80_548 | ~ spl80_1113), inference(avatar_split_clause, [], [f11362, f8245, f4277, f2095, f4265])).
fof(f2095, plain, (spl80_178 <=> (e22 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_178])])).
fof(f11362, plain, ((e22 = h1(e14)) | (~ spl80_178 | ~ spl80_548 | ~ spl80_1113)), inference(backward_demodulation, [], [f11102, f2097])).
fof(f2097, plain, ((e22 = op2(e22, e24)) | ~ spl80_178), inference(avatar_component_clause, [], [f2095])).
fof(f11360, plain, (~ spl80_181 | ~ spl80_184), inference(avatar_contradiction_clause, [], [f11359])).
fof(f11359, plain, ($false | (~ spl80_181 | ~ spl80_184)), inference(subsumption_resolution, [], [f11357, f612])).
fof(f11357, plain, ((e20 = e23) | (~ spl80_181 | ~ spl80_184)), inference(backward_demodulation, [], [f2122, f2110])).
fof(f2122, plain, ((e23 = op2(e22, e23)) | ~ spl80_184), inference(avatar_component_clause, [], [f2120])).
fof(f2120, plain, (spl80_184 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_184])])).
fof(f11309, plain, (~ spl80_226 | ~ spl80_298), inference(avatar_contradiction_clause, [], [f11308])).
fof(f11308, plain, ($false | (~ spl80_226 | ~ spl80_298)), inference(subsumption_resolution, [], [f11307, f610])).
fof(f610, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f10])).
fof(f11307, plain, ((e20 = e21) | (~ spl80_226 | ~ spl80_298)), inference(forward_demodulation, [], [f11297, f2299])).
fof(f2299, plain, ((e20 = op2(e20, e24)) | ~ spl80_226), inference(avatar_component_clause, [], [f2297])).
fof(f2297, plain, (spl80_226 <=> (e20 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_226])])).
fof(f11297, plain, ((e21 = op2(e20, e24)) | ~ spl80_298), inference(backward_demodulation, [], [f2655, f2929])).
fof(f2655, plain, (e21 = op2(h18(e10), e24)), inference(backward_demodulation, [], [f679, f788])).
fof(f679, plain, (e21 = op2(op2(e21, e24), e24)), inference(cnf_transformation, [], [f13])).
fof(f11290, plain, (~ spl80_217 | ~ spl80_121 | ~ spl80_310 | spl80_641), inference(avatar_split_clause, [], [f11284, f4979, f2991, f1806, f2259])).
fof(f2259, plain, (spl80_217 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_217])])).
fof(f1806, plain, (spl80_121 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_121])])).
fof(f4979, plain, (spl80_641 <=> (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_641])])).
fof(f11284, plain, (~ (e21 = op2(e21, e21)) | (~ spl80_121 | ~ spl80_310 | spl80_641)), inference(backward_demodulation, [], [f10719, f2992])).
fof(f10719, plain, (~ (h17(e10) = op2(h17(e10), h17(e10))) | (~ spl80_121 | spl80_641)), inference(backward_demodulation, [], [f4981, f1808])).
fof(f1808, plain, ((e10 = op1(e10, e10)) | ~ spl80_121), inference(avatar_component_clause, [], [f1806])).
fof(f4981, plain, (~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10))) | spl80_641), inference(avatar_component_clause, [], [f4979])).
fof(f11273, plain, (spl80_233 | ~ spl80_346), inference(avatar_split_clause, [], [f11263, f3184, f2326])).
fof(f3184, plain, (spl80_346 <=> (e20 = h15(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_346])])).
fof(f11263, plain, ((e22 = op2(e20, e23)) | ~ spl80_346), inference(backward_demodulation, [], [f2626, f3185])).
fof(f3185, plain, ((e20 = h15(e10)) | ~ spl80_346), inference(avatar_component_clause, [], [f3184])).
fof(f2626, plain, (e22 = op2(h15(e10), e23)), inference(backward_demodulation, [], [f683, f773])).
fof(f773, plain, (op2(e22, e23) = h15(e10)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ((h15(e14) = op2(op2(e22, e23), op2(e23, op2(e22, e23)))) & (h15(e12) = op2(e23, op2(e22, e23))) & (op2(e22, e23) = h15(e10)) & (e22 = h15(e13)) & (e23 = h15(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax30)).
fof(f683, plain, (e22 = op2(op2(e22, e23), e23)), inference(cnf_transformation, [], [f13])).
fof(f11229, plain, (spl80_214 | ~ spl80_406), inference(avatar_split_clause, [], [f11220, f3506, f2246])).
fof(f3506, plain, (spl80_406 <=> (e21 = h11(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_406])])).
fof(f11220, plain, ((e23 = op2(e21, e22)) | ~ spl80_406), inference(backward_demodulation, [], [f2586, f3507])).
fof(f3507, plain, ((e21 = h11(e10)) | ~ spl80_406), inference(avatar_component_clause, [], [f3506])).
fof(f2586, plain, (e23 = op2(h11(e10), e22)), inference(backward_demodulation, [], [f687, f753])).
fof(f753, plain, (op2(e23, e22) = h11(e10)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ((h11(e14) = op2(op2(e23, e22), op2(e22, op2(e23, e22)))) & (h11(e12) = op2(e22, op2(e23, e22))) & (op2(e23, e22) = h11(e10)) & (e23 = h11(e13)) & (e22 = h11(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax26)).
fof(f687, plain, (e23 = op2(op2(e23, e22), e22)), inference(cnf_transformation, [], [f13])).
fof(f11200, plain, ~ spl80_468, inference(avatar_contradiction_clause, [], [f11199])).
fof(f11199, plain, ($false | ~ spl80_468), inference(subsumption_resolution, [], [f11198, f614])).
fof(f614, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f10])).
fof(f11198, plain, ((e21 = e22) | ~ spl80_468), inference(forward_demodulation, [], [f11197, f11187])).
fof(f11187, plain, ((e21 = op2(h6(e12), e22)) | ~ spl80_468), inference(backward_demodulation, [], [f2574, f11184])).
fof(f11184, plain, ((h6(e12) = h10(e10)) | ~ spl80_468), inference(backward_demodulation, [], [f748, f11178])).
fof(f11178, plain, ((op2(e21, e22) = h6(e12)) | ~ spl80_468), inference(backward_demodulation, [], [f2529, f3850])).
fof(f3850, plain, ((e22 = h6(e10)) | ~ spl80_468), inference(avatar_component_clause, [], [f3849])).
fof(f3849, plain, (spl80_468 <=> (e22 = h6(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_468])])).
fof(f2529, plain, (h6(e12) = op2(e21, h6(e10))), inference(backward_demodulation, [], [f729, f728])).
fof(f728, plain, (op2(e22, e21) = h6(e10)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ((h6(e14) = op2(op2(e22, e21), op2(e21, op2(e22, e21)))) & (h6(e12) = op2(e21, op2(e22, e21))) & (op2(e22, e21) = h6(e10)) & (e22 = h6(e13)) & (e21 = h6(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax21)).
fof(f729, plain, (h6(e12) = op2(e21, op2(e22, e21))), inference(cnf_transformation, [], [f21])).
fof(f748, plain, (op2(e21, e22) = h10(e10)), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ((h10(e14) = op2(op2(e21, e22), op2(e22, op2(e21, e22)))) & (h10(e12) = op2(e22, op2(e21, e22))) & (op2(e21, e22) = h10(e10)) & (e21 = h10(e13)) & (e22 = h10(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax25)).
fof(f2574, plain, (e21 = op2(h10(e10), e22)), inference(backward_demodulation, [], [f677, f748])).
fof(f677, plain, (e21 = op2(op2(e21, e22), e22)), inference(cnf_transformation, [], [f13])).
fof(f11197, plain, ((e22 = op2(h6(e12), e22)) | ~ spl80_468), inference(forward_demodulation, [], [f11181, f11184])).
fof(f11181, plain, ((e22 = op2(h10(e10), e22)) | ~ spl80_468), inference(backward_demodulation, [], [f2577, f3850])).
fof(f2577, plain, (e22 = op2(h10(e10), h6(e10))), inference(backward_demodulation, [], [f2521, f748])).
fof(f2521, plain, (e22 = op2(op2(e21, e22), h6(e10))), inference(backward_demodulation, [], [f386, f728])).
fof(f386, plain, (e22 = op2(op2(e21, e22), op2(e22, e21))), inference(cnf_transformation, [], [f6])).
fof(f11172, plain, (spl80_201 | ~ spl80_536 | ~ spl80_892), inference(avatar_split_clause, [], [f11164, f6697, f4213, f2192])).
fof(f2192, plain, (spl80_201 <=> (e20 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_201])])).
fof(f11164, plain, ((e20 = op2(e21, e24)) | (~ spl80_536 | ~ spl80_892)), inference(backward_demodulation, [], [f10509, f4214])).
fof(f10509, plain, ((e20 = op2(h2(e10), e24)) | ~ spl80_892), inference(backward_demodulation, [], [f2564, f6698])).
fof(f2564, plain, (e20 = op2(h2(e10), h9(e10))), inference(backward_demodulation, [], [f2470, f743])).
fof(f2470, plain, (e20 = op2(h2(e10), op2(e20, e22))), inference(backward_demodulation, [], [f377, f708])).
fof(f377, plain, (e20 = op2(op2(e22, e20), op2(e20, e22))), inference(cnf_transformation, [], [f6])).
fof(f11137, plain, ~ spl80_1030, inference(avatar_contradiction_clause, [], [f11136])).
fof(f11136, plain, ($false | ~ spl80_1030), inference(subsumption_resolution, [], [f11135, f613])).
fof(f613, plain, ~ (e20 = e24), inference(cnf_transformation, [], [f10])).
fof(f11135, plain, ((e20 = e24) | ~ spl80_1030), inference(forward_demodulation, [], [f11125, f2645])).
fof(f2645, plain, (e20 = op2(h17(e10), e24)), inference(backward_demodulation, [], [f674, f783])).
fof(f674, plain, (e20 = op2(op2(e20, e24), e24)), inference(cnf_transformation, [], [f13])).
fof(f11125, plain, ((e24 = op2(h17(e10), e24)) | ~ spl80_1030), inference(backward_demodulation, [], [f2648, f7717])).
fof(f7717, plain, ((e24 = h4(e10)) | ~ spl80_1030), inference(avatar_component_clause, [], [f7716])).
fof(f7716, plain, (spl80_1030 <=> (e24 = h4(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_1030])])).
fof(f11098, plain, (spl80_468 | ~ spl80_193), inference(avatar_split_clause, [], [f10648, f2158, f3849])).
fof(f2158, plain, (spl80_193 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_193])])).
fof(f10648, plain, ((e22 = h6(e10)) | ~ spl80_193), inference(backward_demodulation, [], [f728, f2160])).
fof(f2160, plain, ((e22 = op2(e22, e21)) | ~ spl80_193), inference(avatar_component_clause, [], [f2158])).
fof(f11093, plain, (~ spl80_234 | ~ spl80_480), inference(avatar_split_clause, [], [f11092, f3913, f2330])).
fof(f2330, plain, (spl80_234 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_234])])).
fof(f3913, plain, (spl80_480 <=> (e23 = h5(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_480])])).
fof(f11092, plain, (~ (e23 = op2(e20, e23)) | ~ spl80_480), inference(forward_demodulation, [], [f2512, f3914])).
fof(f3914, plain, ((e23 = h5(e10)) | ~ spl80_480), inference(avatar_component_clause, [], [f3913])).
fof(f2512, plain, ~ (op2(e20, e23) = h5(e10)), inference(backward_demodulation, [], [f555, f723])).
fof(f723, plain, (op2(e20, e21) = h5(e10)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ((h5(e14) = op2(op2(e20, e21), op2(e21, op2(e20, e21)))) & (h5(e12) = op2(e21, op2(e20, e21))) & (op2(e20, e21) = h5(e10)) & (e20 = h5(e13)) & (e21 = h5(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax20)).
fof(f555, plain, ~ (op2(e20, e21) = op2(e20, e23)), inference(cnf_transformation, [], [f8])).
fof(f11080, plain, (spl80_142 | ~ spl80_220), inference(avatar_split_clause, [], [f11079, f2271, f1944])).
fof(f1944, plain, (spl80_142 <=> (e21 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_142])])).
fof(f2271, plain, (spl80_220 <=> (e24 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_220])])).
fof(f11079, plain, ((e21 = op2(e24, e21)) | ~ spl80_220), inference(forward_demodulation, [], [f676, f2273])).
fof(f2273, plain, ((e24 = op2(e21, e21)) | ~ spl80_220), inference(avatar_component_clause, [], [f2271])).
fof(f676, plain, (e21 = op2(op2(e21, e21), e21)), inference(cnf_transformation, [], [f13])).
fof(f11078, plain, (~ spl80_218 | ~ spl80_548), inference(avatar_split_clause, [], [f11077, f4277, f2263])).
fof(f2263, plain, (spl80_218 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_218])])).
fof(f11077, plain, (~ (e22 = op2(e21, e21)) | ~ spl80_548), inference(forward_demodulation, [], [f2462, f4278])).
fof(f2462, plain, ~ (op2(e21, e21) = h1(e10)), inference(backward_demodulation, [], [f560, f703])).
fof(f560, plain, ~ (op2(e21, e20) = op2(e21, e21)), inference(cnf_transformation, [], [f8])).
fof(f11076, plain, (~ spl80_219 | ~ spl80_480), inference(avatar_split_clause, [], [f11075, f3913, f2267])).
fof(f2267, plain, (spl80_219 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_219])])).
fof(f11075, plain, (~ (e23 = op2(e21, e21)) | ~ spl80_480), inference(forward_demodulation, [], [f2506, f3914])).
fof(f2506, plain, ~ (op2(e21, e21) = h5(e10)), inference(backward_demodulation, [], [f510, f723])).
fof(f510, plain, ~ (op2(e20, e21) = op2(e21, e21)), inference(cnf_transformation, [], [f8])).
fof(f11065, plain, (spl80_197 | ~ spl80_548), inference(avatar_split_clause, [], [f10524, f4277, f2175])).
fof(f10524, plain, ((e21 = op2(e22, e20)) | ~ spl80_548), inference(backward_demodulation, [], [f2466, f4278])).
fof(f2466, plain, (e21 = op2(h1(e10), e20)), inference(backward_demodulation, [], [f675, f703])).
fof(f675, plain, (e21 = op2(op2(e21, e20), e20)), inference(cnf_transformation, [], [f13])).
fof(f11033, plain, (spl80_181 | ~ spl80_480 | ~ spl80_548), inference(avatar_split_clause, [], [f10547, f4277, f3913, f2108])).
fof(f10547, plain, ((e20 = op2(e22, e23)) | (~ spl80_480 | ~ spl80_548)), inference(backward_demodulation, [], [f10527, f3914])).
fof(f10527, plain, ((e20 = op2(e22, h5(e10))) | ~ spl80_548), inference(backward_demodulation, [], [f2516, f4278])).
fof(f2516, plain, (e20 = op2(h1(e10), h5(e10))), inference(backward_demodulation, [], [f2456, f723])).
fof(f2456, plain, (e20 = op2(h1(e10), op2(e20, e21))), inference(backward_demodulation, [], [f376, f703])).
fof(f376, plain, (e20 = op2(op2(e21, e20), op2(e20, e21))), inference(cnf_transformation, [], [f6])).
fof(f10980, plain, (~ spl80_109 | ~ spl80_119), inference(avatar_split_clause, [], [f10979, f1797, f1755])).
fof(f1755, plain, (spl80_109 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_109])])).
fof(f10979, plain, (~ (e13 = op1(e10, e13)) | ~ spl80_119), inference(forward_demodulation, [], [f455, f1799])).
fof(f455, plain, ~ (op1(e10, e11) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f10972, plain, (~ spl80_105 | ~ spl80_115), inference(avatar_split_clause, [], [f10971, f1780, f1738])).
fof(f1738, plain, (spl80_105 <=> (e14 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_105])])).
fof(f10971, plain, (~ (e14 = op1(e10, e14)) | ~ spl80_115), inference(forward_demodulation, [], [f458, f1782])).
fof(f458, plain, ~ (op1(e10, e12) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f10968, plain, (~ spl80_93 | ~ spl80_98), inference(avatar_split_clause, [], [f10967, f1709, f1688])).
fof(f1688, plain, (spl80_93 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_93])])).
fof(f10967, plain, (~ (e12 = op1(e11, e11)) | ~ spl80_98), inference(forward_demodulation, [], [f460, f1711])).
fof(f460, plain, ~ (op1(e11, e10) = op1(e11, e11)), inference(cnf_transformation, [], [f7])).
fof(f10958, plain, (spl80_85 | ~ spl80_7), inference(avatar_split_clause, [], [f10845, f1327, f1654])).
fof(f10845, plain, ((e14 = op1(e11, e13)) | ~ spl80_7), inference(backward_demodulation, [], [f668, f1329])).
fof(f668, plain, (e14 = op1(op1(e14, e13), e13)), inference(cnf_transformation, [], [f12])).
fof(f10952, plain, (spl80_102 | ~ spl80_76), inference(avatar_split_clause, [], [f10951, f1617, f1726])).
fof(f10951, plain, ((e11 = op1(e10, e14)) | ~ spl80_76), inference(forward_demodulation, [], [f654, f1619])).
fof(f654, plain, (e11 = op1(op1(e11, e14), e14)), inference(cnf_transformation, [], [f12])).
fof(f10922, plain, (spl80_72 | ~ spl80_98), inference(avatar_split_clause, [], [f10386, f1709, f1600])).
fof(f10386, plain, ((e11 = op1(e12, e10)) | ~ spl80_98), inference(forward_demodulation, [], [f650, f1711])).
fof(f650, plain, (e11 = op1(op1(e11, e10), e10)), inference(cnf_transformation, [], [f12])).
fof(f10917, plain, (~ spl80_64 | ~ spl80_89), inference(avatar_split_clause, [], [f10916, f1671, f1566])).
fof(f1566, plain, (spl80_64 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_64])])).
fof(f10916, plain, (~ (e13 = op1(e12, e12)) | ~ spl80_89), inference(forward_demodulation, [], [f424, f1673])).
fof(f424, plain, ~ (op1(e11, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f7])).
fof(f10915, plain, (~ spl80_65 | ~ spl80_115), inference(avatar_split_clause, [], [f10914, f1780, f1570])).
fof(f1570, plain, (spl80_65 <=> (e14 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_65])])).
fof(f10914, plain, (~ (e14 = op1(e12, e12)) | ~ spl80_115), inference(forward_demodulation, [], [f421, f1782])).
fof(f421, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f7])).
fof(f10907, plain, (spl80_108 | ~ spl80_56), inference(avatar_split_clause, [], [f10906, f1533, f1751])).
fof(f10906, plain, ((e12 = op1(e10, e13)) | ~ spl80_56), inference(forward_demodulation, [], [f658, f1535])).
fof(f658, plain, (e12 = op1(op1(e12, e13), e13)), inference(cnf_transformation, [], [f12])).
fof(f10857, plain, (~ spl80_7 | ~ spl80_30 | ~ spl80_76), inference(avatar_contradiction_clause, [], [f10856])).
fof(f10856, plain, ($false | (~ spl80_7 | ~ spl80_30 | ~ spl80_76)), inference(subsumption_resolution, [], [f10855, f602])).
fof(f10855, plain, ((e10 = e13) | (~ spl80_7 | ~ spl80_30 | ~ spl80_76)), inference(forward_demodulation, [], [f10849, f1619])).
fof(f10849, plain, ((e13 = op1(e11, e14)) | (~ spl80_7 | ~ spl80_30)), inference(backward_demodulation, [], [f10813, f1329])).
fof(f10813, plain, ((e13 = op1(op1(e14, e13), e14)) | ~ spl80_30), inference(backward_demodulation, [], [f294, f1425])).
fof(f1425, plain, ((e14 = op1(e13, e14)) | ~ spl80_30), inference(avatar_component_clause, [], [f1423])).
fof(f1423, plain, (spl80_30 <=> (e14 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_30])])).
fof(f294, plain, (e13 = op1(op1(e14, e13), op1(e13, e14))), inference(cnf_transformation, [], [f3])).
fof(f10772, plain, (spl80_57 | ~ spl80_68 | ~ spl80_89), inference(avatar_split_clause, [], [f10767, f1671, f1583, f1537])).
fof(f1537, plain, (spl80_57 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_57])])).
fof(f1583, plain, (spl80_68 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_68])])).
fof(f10767, plain, ((e11 = op1(e12, e13)) | (~ spl80_68 | ~ spl80_89)), inference(backward_demodulation, [], [f10375, f1585])).
fof(f1585, plain, ((e12 = op1(e12, e11)) | ~ spl80_68), inference(avatar_component_clause, [], [f1583])).
fof(f10375, plain, ((e11 = op1(op1(e12, e11), e13)) | ~ spl80_89), inference(forward_demodulation, [], [f282, f1673])).
fof(f282, plain, (e11 = op1(op1(e12, e11), op1(e11, e12))), inference(cnf_transformation, [], [f3])).
fof(f10759, plain, ~ spl80_77, inference(avatar_contradiction_clause, [], [f10758])).
fof(f10758, plain, ($false | ~ spl80_77), inference(subsumption_resolution, [], [f10757, f606])).
fof(f606, plain, ~ (e11 = e14), inference(cnf_transformation, [], [f9])).
fof(f10757, plain, ((e11 = e14) | ~ spl80_77), inference(backward_demodulation, [], [f666, f10749])).
fof(f10749, plain, ((e11 = op1(op1(e14, e11), e11)) | ~ spl80_77), inference(backward_demodulation, [], [f284, f1623])).
fof(f1623, plain, ((e11 = op1(e11, e14)) | ~ spl80_77), inference(avatar_component_clause, [], [f1621])).
fof(f1621, plain, (spl80_77 <=> (e11 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_77])])).
fof(f284, plain, (e11 = op1(op1(e14, e11), op1(e11, e14))), inference(cnf_transformation, [], [f3])).
fof(f666, plain, (e14 = op1(op1(e14, e11), e11)), inference(cnf_transformation, [], [f12])).
fof(f10631, plain, (~ spl80_223 | ~ spl80_225), inference(avatar_contradiction_clause, [], [f10630])).
fof(f10630, plain, ($false | (~ spl80_223 | ~ spl80_225)), inference(subsumption_resolution, [], [f10629, f618])).
fof(f618, plain, ~ (e22 = e24), inference(cnf_transformation, [], [f10])).
fof(f10629, plain, ((e22 = e24) | (~ spl80_223 | ~ spl80_225)), inference(forward_demodulation, [], [f2294, f2286])).
fof(f2294, plain, ((e24 = op2(e21, e20)) | ~ spl80_225), inference(avatar_component_clause, [], [f2292])).
fof(f2292, plain, (spl80_225 <=> (e24 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_225])])).
fof(f10623, plain, (~ spl80_370 | ~ spl80_228), inference(avatar_split_clause, [], [f10620, f2305, f3313])).
fof(f2305, plain, (spl80_228 <=> (e22 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_228])])).
fof(f10620, plain, (~ (e22 = h13(e10)) | ~ spl80_228), inference(backward_demodulation, [], [f2604, f2307])).
fof(f2307, plain, ((e22 = op2(e20, e24)) | ~ spl80_228), inference(avatar_component_clause, [], [f2305])).
fof(f2604, plain, ~ (op2(e20, e24) = h13(e10)), inference(backward_demodulation, [], [f559, f763])).
fof(f559, plain, ~ (op2(e20, e23) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f10614, plain, (spl80_216 | ~ spl80_247), inference(avatar_contradiction_clause, [], [f10613])).
fof(f10613, plain, ($false | (spl80_216 | ~ spl80_247)), inference(subsumption_resolution, [], [f10605, f2256])).
fof(f2256, plain, (~ (e20 = op2(e21, e21)) | spl80_216), inference(avatar_component_clause, [], [f2255])).
fof(f2255, plain, (spl80_216 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_216])])).
fof(f10605, plain, ((e20 = op2(e21, e21)) | ~ spl80_247), inference(backward_demodulation, [], [f375, f2387])).
fof(f2387, plain, ((op2(e20, e20) = e21) | ~ spl80_247), inference(avatar_component_clause, [], [f2385])).
fof(f2385, plain, (spl80_247 <=> (op2(e20, e20) = e21)), introduced(avatar_definition, [new_symbols(naming, [spl80_247])])).
fof(f375, plain, (e20 = op2(op2(e20, e20), op2(e20, e20))), inference(cnf_transformation, [], [f6])).
fof(f10552, plain, (~ spl80_229 | ~ spl80_480), inference(avatar_split_clause, [], [f10541, f3913, f2309])).
fof(f2309, plain, (spl80_229 <=> (e23 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_229])])).
fof(f10541, plain, (~ (e23 = op2(e20, e24)) | ~ spl80_480), inference(backward_demodulation, [], [f2513, f3914])).
fof(f2513, plain, ~ (op2(e20, e24) = h5(e10)), inference(backward_demodulation, [], [f556, f723])).
fof(f556, plain, ~ (op2(e20, e21) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f10551, plain, (~ spl80_249 | ~ spl80_480), inference(avatar_split_clause, [], [f10540, f3913, f2393])).
fof(f2393, plain, (spl80_249 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl80_249])])).
fof(f10540, plain, (~ (op2(e20, e20) = e23) | ~ spl80_480), inference(backward_demodulation, [], [f2510, f3914])).
fof(f2510, plain, ~ (op2(e20, e20) = h5(e10)), inference(backward_demodulation, [], [f550, f723])).
fof(f550, plain, ~ (op2(e20, e20) = op2(e20, e21)), inference(cnf_transformation, [], [f8])).
fof(f10537, plain, (spl80_1113 | ~ spl80_240 | ~ spl80_548), inference(avatar_split_clause, [], [f10536, f4277, f2355, f8245])).
fof(f10536, plain, ((e24 = h1(e12)) | (~ spl80_240 | ~ spl80_548)), inference(forward_demodulation, [], [f10525, f2357])).
fof(f10525, plain, ((op2(e20, e22) = h1(e12)) | ~ spl80_548), inference(backward_demodulation, [], [f2467, f4278])).
fof(f2467, plain, (h1(e12) = op2(e20, h1(e10))), inference(backward_demodulation, [], [f704, f703])).
fof(f10530, plain, (~ spl80_198 | ~ spl80_548), inference(avatar_split_clause, [], [f10520, f4277, f2179])).
fof(f2179, plain, (spl80_198 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_198])])).
fof(f10520, plain, (~ (e22 = op2(e22, e20)) | ~ spl80_548), inference(backward_demodulation, [], [f2459, f4278])).
fof(f2459, plain, ~ (op2(e22, e20) = h1(e10)), inference(backward_demodulation, [], [f504, f703])).
fof(f504, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f8])).
fof(f10529, plain, (~ spl80_248 | ~ spl80_548), inference(avatar_split_clause, [], [f10519, f4277, f2389])).
fof(f2389, plain, (spl80_248 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl80_248])])).
fof(f10519, plain, (~ (op2(e20, e20) = e22) | ~ spl80_548), inference(backward_demodulation, [], [f2458, f4278])).
fof(f2458, plain, ~ (op2(e20, e20) = h1(e10)), inference(backward_demodulation, [], [f500, f703])).
fof(f500, plain, ~ (op2(e20, e20) = op2(e21, e20)), inference(cnf_transformation, [], [f8])).
fof(f10515, plain, (spl80_136 | ~ spl80_892), inference(avatar_split_clause, [], [f10507, f6697, f1919])).
fof(f10507, plain, ((e20 = op2(e24, e22)) | ~ spl80_892), inference(backward_demodulation, [], [f2562, f6698])).
fof(f2562, plain, (e20 = op2(h9(e10), e22)), inference(backward_demodulation, [], [f672, f743])).
fof(f672, plain, (e20 = op2(op2(e20, e22), e22)), inference(cnf_transformation, [], [f13])).
fof(f10514, plain, (~ spl80_250 | ~ spl80_892), inference(avatar_split_clause, [], [f10505, f6697, f2397])).
fof(f2397, plain, (spl80_250 <=> (op2(e20, e20) = e24)), introduced(avatar_definition, [new_symbols(naming, [spl80_250])])).
fof(f10505, plain, (~ (op2(e20, e20) = e24) | ~ spl80_892), inference(backward_demodulation, [], [f2559, f6698])).
fof(f2559, plain, ~ (op2(e20, e20) = h9(e10)), inference(backward_demodulation, [], [f551, f743])).
fof(f551, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f8])).
fof(f10453, plain, (spl80_406 | ~ spl80_162), inference(avatar_split_clause, [], [f10452, f2028, f3506])).
fof(f10452, plain, ((e21 = h11(e10)) | ~ spl80_162), inference(forward_demodulation, [], [f753, f2030])).
fof(f10422, plain, (~ spl80_40 | ~ spl80_115), inference(avatar_split_clause, [], [f10421, f1780, f1465])).
fof(f1465, plain, (spl80_40 <=> (e14 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_40])])).
fof(f10421, plain, (~ (e14 = op1(e13, e12)) | ~ spl80_115), inference(forward_demodulation, [], [f422, f1782])).
fof(f422, plain, ~ (op1(e10, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f10417, plain, (spl80_11 | ~ spl80_115), inference(avatar_split_clause, [], [f10416, f1780, f1344])).
fof(f10416, plain, ((e10 = op1(e14, e12)) | ~ spl80_115), inference(forward_demodulation, [], [f647, f1782])).
fof(f647, plain, (e10 = op1(op1(e10, e12), e12)), inference(cnf_transformation, [], [f12])).
fof(f10383, plain, (spl80_37 | ~ spl80_98 | ~ spl80_119), inference(avatar_split_clause, [], [f10382, f1797, f1709, f1453])).
fof(f10382, plain, ((e11 = op1(e13, e12)) | (~ spl80_98 | ~ spl80_119)), inference(forward_demodulation, [], [f10381, f1799])).
fof(f10381, plain, ((e11 = op1(op1(e10, e11), e12)) | ~ spl80_98), inference(forward_demodulation, [], [f280, f1711])).
fof(f280, plain, (e11 = op1(op1(e10, e11), op1(e11, e10))), inference(cnf_transformation, [], [f3])).
fof(f10380, plain, (spl80_56 | ~ spl80_98 | ~ spl80_119), inference(avatar_split_clause, [], [f10379, f1797, f1709, f1533])).
fof(f10379, plain, ((e10 = op1(e12, e13)) | (~ spl80_98 | ~ spl80_119)), inference(forward_demodulation, [], [f10378, f1711])).
fof(f10378, plain, ((e10 = op1(op1(e11, e10), e13)) | ~ spl80_119), inference(forward_demodulation, [], [f276, f1799])).
fof(f276, plain, (e10 = op1(op1(e11, e10), op1(e10, e11))), inference(cnf_transformation, [], [f3])).
fof(f10346, plain, (~ spl80_69 | ~ spl80_119), inference(avatar_split_clause, [], [f10345, f1797, f1587])).
fof(f1587, plain, (spl80_69 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_69])])).
fof(f10345, plain, (~ (e13 = op1(e12, e11)) | ~ spl80_119), inference(forward_demodulation, [], [f411, f1799])).
fof(f411, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f7])).
fof(f10325, plain, (~ spl80_59 | ~ spl80_54), inference(avatar_split_clause, [], [f10324, f1524, f1545])).
fof(f1545, plain, (spl80_59 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_59])])).
fof(f10324, plain, (~ (e13 = op1(e12, e13)) | ~ spl80_54), inference(forward_demodulation, [], [f479, f1526])).
fof(f479, plain, ~ (op1(e12, e13) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f10254, plain, (~ spl80_21 | ~ spl80_24), inference(avatar_contradiction_clause, [], [f10253])).
fof(f10253, plain, ($false | (~ spl80_21 | ~ spl80_24)), inference(subsumption_resolution, [], [f10252, f602])).
fof(f10252, plain, ((e10 = e13) | (~ spl80_21 | ~ spl80_24)), inference(backward_demodulation, [], [f1400, f1388])).
fof(f1388, plain, ((e10 = op1(e14, e10)) | ~ spl80_21), inference(avatar_component_clause, [], [f1386])).
fof(f1386, plain, (spl80_21 <=> (e10 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_21])])).
fof(f10160, plain, (~ spl80_104 | ~ spl80_119), inference(avatar_split_clause, [], [f10155, f1797, f1734])).
fof(f1734, plain, (spl80_104 <=> (e13 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_104])])).
fof(f10155, plain, (~ (e13 = op1(e10, e14)) | ~ spl80_119), inference(backward_demodulation, [], [f456, f1799])).
fof(f456, plain, ~ (op1(e10, e11) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f10158, plain, (~ spl80_94 | ~ spl80_119), inference(avatar_split_clause, [], [f10153, f1797, f1692])).
fof(f1692, plain, (spl80_94 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_94])])).
fof(f10153, plain, (~ (e13 = op1(e11, e11)) | ~ spl80_119), inference(backward_demodulation, [], [f410, f1799])).
fof(f410, plain, ~ (op1(e10, e11) = op1(e11, e11)), inference(cnf_transformation, [], [f7])).
fof(f10104, plain, (~ spl80_127 | ~ spl80_130), inference(avatar_contradiction_clause, [], [f10103])).
fof(f10103, plain, ($false | (~ spl80_127 | ~ spl80_130)), inference(subsumption_resolution, [], [f10102, f616])).
fof(f10102, plain, ((e21 = e24) | (~ spl80_127 | ~ spl80_130)), inference(backward_demodulation, [], [f1895, f1883])).
fof(f1883, plain, ((e21 = op2(e24, e24)) | ~ spl80_127), inference(avatar_component_clause, [], [f1881])).
fof(f1881, plain, (spl80_127 <=> (e21 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_127])])).
fof(f10098, plain, (spl80_450 | ~ spl80_143), inference(avatar_split_clause, [], [f10093, f1948, f3741])).
fof(f10093, plain, ((e22 = h8(e10)) | ~ spl80_143), inference(backward_demodulation, [], [f738, f1950])).
fof(f1950, plain, ((e22 = op2(e24, e21)) | ~ spl80_143), inference(avatar_component_clause, [], [f1948])).
fof(f10035, plain, (~ spl80_402 | ~ spl80_188), inference(avatar_split_clause, [], [f10027, f2137, f3485])).
fof(f3485, plain, (spl80_402 <=> (e22 = h11(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_402])])).
fof(f10027, plain, (~ (e22 = h11(e10)) | ~ spl80_188), inference(backward_demodulation, [], [f2582, f2139])).
fof(f2582, plain, ~ (op2(e22, e22) = h11(e10)), inference(backward_demodulation, [], [f527, f753])).
fof(f527, plain, ~ (op2(e22, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f8])).
fof(f10019, plain, (~ spl80_1086 | ~ spl80_195), inference(avatar_split_clause, [], [f10016, f2166, f8088])).
fof(f8088, plain, (spl80_1086 <=> (e24 = h2(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_1086])])).
fof(f10016, plain, (~ (e24 = h2(e10)) | ~ spl80_195), inference(backward_demodulation, [], [f2475, f2168])).
fof(f2475, plain, ~ (op2(e22, e21) = h2(e10)), inference(backward_demodulation, [], [f570, f708])).
fof(f570, plain, ~ (op2(e22, e20) = op2(e22, e21)), inference(cnf_transformation, [], [f8])).
fof(f10018, plain, (spl80_974 | ~ spl80_195), inference(avatar_split_clause, [], [f10015, f2166, f7320])).
fof(f7320, plain, (spl80_974 <=> (e24 = h6(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_974])])).
fof(f10015, plain, ((e24 = h6(e10)) | ~ spl80_195), inference(backward_demodulation, [], [f728, f2168])).
fof(f10013, plain, (spl80_1086 | ~ spl80_200), inference(avatar_split_clause, [], [f10011, f2187, f8088])).
fof(f2187, plain, (spl80_200 <=> (e24 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_200])])).
fof(f10011, plain, ((e24 = h2(e10)) | ~ spl80_200), inference(backward_demodulation, [], [f708, f2189])).
fof(f2189, plain, ((e24 = op2(e22, e20)) | ~ spl80_200), inference(avatar_component_clause, [], [f2187])).
fof(f10002, plain, (spl80_548 | ~ spl80_223), inference(avatar_split_clause, [], [f10001, f2284, f4277])).
fof(f10001, plain, ((e22 = h1(e10)) | ~ spl80_223), inference(backward_demodulation, [], [f703, f2286])).
fof(f9995, plain, (spl80_892 | ~ spl80_240), inference(avatar_split_clause, [], [f9993, f2355, f6697])).
fof(f9993, plain, ((e24 = h9(e10)) | ~ spl80_240), inference(backward_demodulation, [], [f743, f2357])).
fof(f9992, plain, (spl80_480 | ~ spl80_244), inference(avatar_split_clause, [], [f9991, f2372, f3913])).
fof(f9991, plain, ((e23 = h5(e10)) | ~ spl80_244), inference(backward_demodulation, [], [f723, f2374])).
fof(f9990, plain, (spl80_370 | ~ spl80_233), inference(avatar_split_clause, [], [f9989, f2326, f3313])).
fof(f9989, plain, ((e22 = h13(e10)) | ~ spl80_233), inference(forward_demodulation, [], [f763, f2328])).
fof(f9958, plain, (spl80_298 | ~ spl80_201), inference(avatar_split_clause, [], [f9957, f2192, f2928])).
fof(f9957, plain, ((e20 = h18(e10)) | ~ spl80_201), inference(forward_demodulation, [], [f788, f2194])).
fof(f2194, plain, ((e20 = op2(e21, e24)) | ~ spl80_201), inference(avatar_component_clause, [], [f2192])).
fof(f9948, plain, (~ spl80_370 | ~ spl80_183), inference(avatar_split_clause, [], [f9947, f2116, f3313])).
fof(f2116, plain, (spl80_183 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_183])])).
fof(f9947, plain, (~ (e22 = h13(e10)) | ~ spl80_183), inference(forward_demodulation, [], [f2600, f2118])).
fof(f2118, plain, ((e22 = op2(e22, e23)) | ~ spl80_183), inference(avatar_component_clause, [], [f2116])).
fof(f2600, plain, ~ (op2(e22, e23) = h13(e10)), inference(backward_demodulation, [], [f531, f763])).
fof(f531, plain, ~ (op2(e20, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f8])).
fof(f9928, plain, (~ spl80_548 | ~ spl80_173), inference(avatar_split_clause, [], [f9927, f2074, f4277])).
fof(f2074, plain, (spl80_173 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_173])])).
fof(f9927, plain, (~ (e22 = h1(e10)) | ~ spl80_173), inference(forward_demodulation, [], [f2460, f2076])).
fof(f2076, plain, ((e22 = op2(e23, e20)) | ~ spl80_173), inference(avatar_component_clause, [], [f2074])).
fof(f2460, plain, ~ (op2(e23, e20) = h1(e10)), inference(backward_demodulation, [], [f505, f703])).
fof(f505, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f8])).
fof(f9924, plain, (spl80_402 | ~ spl80_163), inference(avatar_split_clause, [], [f9923, f2032, f3485])).
fof(f2032, plain, (spl80_163 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_163])])).
fof(f9923, plain, ((e22 = h11(e10)) | ~ spl80_163), inference(forward_demodulation, [], [f753, f2034])).
fof(f2034, plain, ((e22 = op2(e23, e22)) | ~ spl80_163), inference(avatar_component_clause, [], [f2032])).
fof(f9914, plain, (spl80_496 | ~ spl80_149), inference(avatar_split_clause, [], [f9909, f1973, f3999])).
fof(f9909, plain, ((e23 = h4(e10)) | ~ spl80_149), inference(backward_demodulation, [], [f718, f1975])).
fof(f1975, plain, ((e23 = op2(e24, e20)) | ~ spl80_149), inference(avatar_component_clause, [], [f1973])).
fof(f9892, plain, (spl80_394 | ~ spl80_136), inference(avatar_split_clause, [], [f9891, f1919, f3441])).
fof(f9891, plain, ((e20 = h12(e10)) | ~ spl80_136), inference(forward_demodulation, [], [f758, f1921])).
fof(f9862, plain, ~ spl80_97, inference(avatar_contradiction_clause, [], [f9861])).
fof(f9861, plain, ($false | ~ spl80_97), inference(subsumption_resolution, [], [f9860, f600])).
fof(f600, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f9])).
fof(f9860, plain, ((e10 = e11) | ~ spl80_97), inference(forward_demodulation, [], [f9859, f646])).
fof(f646, plain, (e10 = op1(op1(e10, e11), e11)), inference(cnf_transformation, [], [f12])).
fof(f9859, plain, ((e11 = op1(op1(e10, e11), e11)) | ~ spl80_97), inference(forward_demodulation, [], [f280, f1707])).
fof(f1707, plain, ((e11 = op1(e11, e10)) | ~ spl80_97), inference(avatar_component_clause, [], [f1705])).
fof(f1705, plain, (spl80_97 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_97])])).
fof(f9850, plain, (~ spl80_102 | ~ spl80_2), inference(avatar_split_clause, [], [f9222, f1306, f1726])).
fof(f1306, plain, (spl80_2 <=> (e11 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_2])])).
fof(f9222, plain, (~ (e11 = op1(e10, e14)) | ~ spl80_2), inference(forward_demodulation, [], [f443, f1308])).
fof(f1308, plain, ((e11 = op1(e14, e14)) | ~ spl80_2), inference(avatar_component_clause, [], [f1306])).
fof(f443, plain, ~ (op1(e10, e14) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f9845, plain, (spl80_93 | ~ spl80_67), inference(avatar_split_clause, [], [f9671, f1579, f1688])).
fof(f1579, plain, (spl80_67 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_67])])).
fof(f9671, plain, ((e12 = op1(e11, e11)) | ~ spl80_67), inference(backward_demodulation, [], [f656, f1581])).
fof(f1581, plain, ((e11 = op1(e12, e11)) | ~ spl80_67), inference(avatar_component_clause, [], [f1579])).
fof(f9844, plain, (spl80_94 | ~ spl80_32), inference(avatar_split_clause, [], [f9747, f1432, f1692])).
fof(f1432, plain, (spl80_32 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_32])])).
fof(f9747, plain, ((e13 = op1(e11, e11)) | ~ spl80_32), inference(backward_demodulation, [], [f293, f1434])).
fof(f1434, plain, ((e11 = op1(e13, e13)) | ~ spl80_32), inference(avatar_component_clause, [], [f1432])).
fof(f293, plain, (e13 = op1(op1(e13, e13), op1(e13, e13))), inference(cnf_transformation, [], [f3])).
fof(f9726, plain, (~ spl80_41 | ~ spl80_45), inference(avatar_contradiction_clause, [], [f9725])).
fof(f9725, plain, ($false | (~ spl80_41 | ~ spl80_45)), inference(subsumption_resolution, [], [f9724, f603])).
fof(f603, plain, ~ (e10 = e14), inference(cnf_transformation, [], [f9])).
fof(f9724, plain, ((e10 = e14) | (~ spl80_41 | ~ spl80_45)), inference(forward_demodulation, [], [f1488, f1472])).
fof(f1488, plain, ((e14 = op1(e13, e11)) | ~ spl80_45), inference(avatar_component_clause, [], [f1486])).
fof(f1486, plain, (spl80_45 <=> (e14 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_45])])).
fof(f9660, plain, (~ spl80_57 | ~ spl80_72), inference(avatar_split_clause, [], [f9653, f1600, f1537])).
fof(f9653, plain, (~ (e11 = op1(e12, e13)) | ~ spl80_72), inference(backward_demodulation, [], [f472, f1602])).
fof(f472, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f7])).
fof(f9659, plain, (~ spl80_62 | ~ spl80_72), inference(avatar_split_clause, [], [f9652, f1600, f1558])).
fof(f1558, plain, (spl80_62 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_62])])).
fof(f9652, plain, (~ (e11 = op1(e12, e12)) | ~ spl80_72), inference(backward_demodulation, [], [f471, f1602])).
fof(f471, plain, ~ (op1(e12, e10) = op1(e12, e12)), inference(cnf_transformation, [], [f7])).
fof(f9646, plain, (spl80_32 | ~ spl80_84), inference(avatar_split_clause, [], [f9641, f1650, f1432])).
fof(f1650, plain, (spl80_84 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_84])])).
fof(f9641, plain, ((e11 = op1(e13, e13)) | ~ spl80_84), inference(backward_demodulation, [], [f653, f1652])).
fof(f1652, plain, ((e13 = op1(e11, e13)) | ~ spl80_84), inference(avatar_component_clause, [], [f1650])).
fof(f653, plain, (e11 = op1(op1(e11, e13), e13)), inference(cnf_transformation, [], [f12])).
fof(f9555, plain, (~ spl80_136 | ~ spl80_137), inference(avatar_contradiction_clause, [], [f9554])).
fof(f9554, plain, ($false | (~ spl80_136 | ~ spl80_137)), inference(subsumption_resolution, [], [f9553, f610])).
fof(f9553, plain, ((e20 = e21) | (~ spl80_136 | ~ spl80_137)), inference(backward_demodulation, [], [f1925, f1921])).
fof(f1925, plain, ((e21 = op2(e24, e22)) | ~ spl80_137), inference(avatar_component_clause, [], [f1923])).
fof(f1923, plain, (spl80_137 <=> (e21 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_137])])).
fof(f9511, plain, (~ spl80_166 | ~ spl80_168), inference(avatar_contradiction_clause, [], [f9510])).
fof(f9510, plain, ($false | (~ spl80_166 | ~ spl80_168)), inference(subsumption_resolution, [], [f9509, f611])).
fof(f611, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f10])).
fof(f9509, plain, ((e20 = e22) | (~ spl80_166 | ~ spl80_168)), inference(forward_demodulation, [], [f2055, f2047])).
fof(f2055, plain, ((e22 = op2(e23, e21)) | ~ spl80_168), inference(avatar_component_clause, [], [f2053])).
fof(f2053, plain, (spl80_168 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_168])])).
fof(f9500, plain, (spl80_1058 | ~ spl80_175), inference(avatar_split_clause, [], [f9497, f2082, f7902])).
fof(f9497, plain, ((e24 = h3(e10)) | ~ spl80_175), inference(backward_demodulation, [], [f713, f2084])).
fof(f9471, plain, (~ spl80_536 | ~ spl80_187), inference(avatar_split_clause, [], [f9457, f2133, f4213])).
fof(f2133, plain, (spl80_187 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_187])])).
fof(f9457, plain, (~ (e21 = h2(e10)) | ~ spl80_187), inference(backward_demodulation, [], [f2476, f2135])).
fof(f2135, plain, ((e21 = op2(e22, e22)) | ~ spl80_187), inference(avatar_component_clause, [], [f2133])).
fof(f2476, plain, ~ (op2(e22, e22) = h2(e10)), inference(backward_demodulation, [], [f571, f708])).
fof(f571, plain, ~ (op2(e22, e20) = op2(e22, e22)), inference(cnf_transformation, [], [f8])).
fof(f9447, plain, (spl80_536 | ~ spl80_197), inference(avatar_split_clause, [], [f9445, f2175, f4213])).
fof(f9445, plain, ((e21 = h2(e10)) | ~ spl80_197), inference(backward_demodulation, [], [f708, f2177])).
fof(f9433, plain, (spl80_752 | ~ spl80_210), inference(avatar_split_clause, [], [f9429, f2229, f5743])).
fof(f5743, plain, (spl80_752 <=> (e24 = h14(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_752])])).
fof(f9429, plain, ((e24 = h14(e10)) | ~ spl80_210), inference(backward_demodulation, [], [f768, f2231])).
fof(f768, plain, (op2(e21, e23) = h14(e10)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ((h14(e14) = op2(op2(e21, e23), op2(e23, op2(e21, e23)))) & (h14(e12) = op2(e23, op2(e21, e23))) & (op2(e21, e23) = h14(e10)) & (e21 = h14(e13)) & (e23 = h14(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax29)).
fof(f9426, plain, (spl80_422 | ~ spl80_212), inference(avatar_split_clause, [], [f9423, f2238, f3592])).
fof(f3592, plain, (spl80_422 <=> (e21 = h10(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_422])])).
fof(f2238, plain, (spl80_212 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_212])])).
fof(f9423, plain, ((e21 = h10(e10)) | ~ spl80_212), inference(backward_demodulation, [], [f748, f2240])).
fof(f2240, plain, ((e21 = op2(e21, e22)) | ~ spl80_212), inference(avatar_component_clause, [], [f2238])).
fof(f9420, plain, (~ spl80_422 | ~ spl80_217), inference(avatar_split_clause, [], [f9413, f2259, f3592])).
fof(f9413, plain, (~ (e21 = h10(e10)) | ~ spl80_217), inference(backward_demodulation, [], [f2571, f2261])).
fof(f2261, plain, ((e21 = op2(e21, e21)) | ~ spl80_217), inference(avatar_component_clause, [], [f2259])).
fof(f2571, plain, ~ (op2(e21, e21) = h10(e10)), inference(backward_demodulation, [], [f564, f748])).
fof(f564, plain, ~ (op2(e21, e21) = op2(e21, e22)), inference(cnf_transformation, [], [f8])).
fof(f9392, plain, (spl80_438 | ~ spl80_237), inference(avatar_split_clause, [], [f9390, f2343, f3677])).
fof(f3677, plain, (spl80_438 <=> (e21 = h9(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_438])])).
fof(f2343, plain, (spl80_237 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_237])])).
fof(f9390, plain, ((e21 = h9(e10)) | ~ spl80_237), inference(backward_demodulation, [], [f743, f2345])).
fof(f2345, plain, ((e21 = op2(e20, e22)) | ~ spl80_237), inference(avatar_component_clause, [], [f2343])).
fof(f9385, plain, (spl80_240 | ~ spl80_460), inference(avatar_split_clause, [], [f9375, f3806, f2355])).
fof(f3806, plain, (spl80_460 <=> (e20 = h7(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_460])])).
fof(f9375, plain, ((e24 = op2(e20, e22)) | ~ spl80_460), inference(backward_demodulation, [], [f2543, f3807])).
fof(f3807, plain, ((e20 = h7(e10)) | ~ spl80_460), inference(avatar_component_clause, [], [f3806])).
fof(f2543, plain, (e24 = op2(h7(e10), e22)), inference(backward_demodulation, [], [f2453, f733])).
fof(f733, plain, (op2(e23, e21) = h7(e10)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ((op2(op2(e23, e21), op2(e21, op2(e23, e21))) = h7(e14)) & (op2(e21, op2(e23, e21)) = h7(e12)) & (op2(e23, e21) = h7(e10)) & (e23 = h7(e13)) & (e21 = h7(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax22)).
fof(f2453, plain, (e24 = op2(op2(e23, e21), e22)), inference(backward_demodulation, [], [f700, f699])).
fof(f699, plain, (e22 = op2(e21, op2(e23, e21))), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((e24 = op2(op2(e23, e21), op2(e21, op2(e23, e21)))) & (e22 = op2(e21, op2(e23, e21))) & (e20 = op2(e23, e21))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax15)).
fof(f700, plain, (e24 = op2(op2(e23, e21), op2(e21, op2(e23, e21)))), inference(cnf_transformation, [], [f15])).
fof(f9384, plain, (spl80_223 | ~ spl80_460), inference(avatar_split_clause, [], [f9374, f3806, f2284])).
fof(f9374, plain, ((e22 = op2(e21, e20)) | ~ spl80_460), inference(backward_demodulation, [], [f2542, f3807])).
fof(f2542, plain, (e22 = op2(e21, h7(e10))), inference(backward_demodulation, [], [f699, f733])).
fof(f9383, plain, (spl80_244 | ~ spl80_460), inference(avatar_split_clause, [], [f9373, f3806, f2372])).
fof(f9373, plain, ((e23 = op2(e20, e21)) | ~ spl80_460), inference(backward_demodulation, [], [f2541, f3807])).
fof(f2541, plain, (e23 = op2(h7(e10), e21)), inference(backward_demodulation, [], [f686, f733])).
fof(f686, plain, (e23 = op2(op2(e23, e21), e21)), inference(cnf_transformation, [], [f13])).
fof(f9380, plain, (~ spl80_216 | ~ spl80_460), inference(avatar_split_clause, [], [f9370, f3806, f2255])).
fof(f9370, plain, (~ (e20 = op2(e21, e21)) | ~ spl80_460), inference(backward_demodulation, [], [f2536, f3807])).
fof(f2536, plain, ~ (op2(e21, e21) = h7(e10)), inference(backward_demodulation, [], [f515, f733])).
fof(f515, plain, ~ (op2(e21, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f8])).
fof(f9369, plain, (spl80_310 | ~ spl80_227), inference(avatar_split_clause, [], [f9368, f2301, f2991])).
fof(f2301, plain, (spl80_227 <=> (e21 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_227])])).
fof(f9368, plain, ((e21 = h17(e10)) | ~ spl80_227), inference(forward_demodulation, [], [f783, f2303])).
fof(f2303, plain, ((e21 = op2(e20, e24)) | ~ spl80_227), inference(avatar_component_clause, [], [f2301])).
fof(f9365, plain, (~ spl80_438 | ~ spl80_227), inference(avatar_split_clause, [], [f9364, f2301, f3677])).
fof(f9364, plain, (~ (e21 = h9(e10)) | ~ spl80_227), inference(forward_demodulation, [], [f2561, f2303])).
fof(f2561, plain, ~ (op2(e20, e24) = h9(e10)), inference(backward_demodulation, [], [f558, f743])).
fof(f558, plain, ~ (op2(e20, e22) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f9253, plain, (spl80_119 | ~ spl80_41), inference(avatar_split_clause, [], [f8928, f1470, f1797])).
fof(f8928, plain, ((e13 = op1(e10, e11)) | ~ spl80_41), inference(backward_demodulation, [], [f661, f1472])).
fof(f661, plain, (e13 = op1(op1(e13, e11), e11)), inference(cnf_transformation, [], [f12])).
fof(f9245, plain, (spl80_115 | ~ spl80_41), inference(avatar_split_clause, [], [f8930, f1470, f1780])).
fof(f8930, plain, ((e14 = op1(e10, e12)) | ~ spl80_41), inference(backward_demodulation, [], [f2451, f1472])).
fof(f2451, plain, (e14 = op1(op1(e13, e11), e12)), inference(backward_demodulation, [], [f697, f696])).
fof(f696, plain, (e12 = op1(e11, op1(e13, e11))), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((e14 = op1(op1(e13, e11), op1(e11, op1(e13, e11)))) & (e12 = op1(e11, op1(e13, e11))) & (e10 = op1(e13, e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax14)).
fof(f697, plain, (e14 = op1(op1(e13, e11), op1(e11, op1(e13, e11)))), inference(cnf_transformation, [], [f14])).
fof(f9242, plain, (~ spl80_107 | ~ spl80_102), inference(avatar_split_clause, [], [f9241, f1726, f1747])).
fof(f1747, plain, (spl80_107 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_107])])).
fof(f9241, plain, (~ (e11 = op1(e10, e13)) | ~ spl80_102), inference(forward_demodulation, [], [f459, f1728])).
fof(f459, plain, ~ (op1(e10, e13) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f9210, plain, (spl80_98 | ~ spl80_41), inference(avatar_split_clause, [], [f8929, f1470, f1709])).
fof(f8929, plain, ((e12 = op1(e11, e10)) | ~ spl80_41), inference(backward_demodulation, [], [f696, f1472])).
fof(f9209, plain, (~ spl80_91 | ~ spl80_41), inference(avatar_split_clause, [], [f9208, f1470, f1680])).
fof(f1680, plain, (spl80_91 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_91])])).
fof(f9208, plain, (~ (e10 = op1(e11, e11)) | ~ spl80_41), inference(forward_demodulation, [], [f415, f1472])).
fof(f415, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f9186, plain, (~ spl80_66 | ~ spl80_41), inference(avatar_split_clause, [], [f9185, f1470, f1575])).
fof(f1575, plain, (spl80_66 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_66])])).
fof(f9185, plain, (~ (e10 = op1(e12, e11)) | ~ spl80_41), inference(forward_demodulation, [], [f417, f1472])).
fof(f417, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f9153, plain, (~ spl80_46 | ~ spl80_41), inference(avatar_split_clause, [], [f9152, f1470, f1491])).
fof(f1491, plain, (spl80_46 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_46])])).
fof(f9152, plain, (~ (e10 = op1(e13, e10)) | ~ spl80_41), inference(forward_demodulation, [], [f480, f1472])).
fof(f480, plain, ~ (op1(e13, e10) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f8993, plain, (~ spl80_28 | ~ spl80_29), inference(avatar_contradiction_clause, [], [f8992])).
fof(f8992, plain, ($false | (~ spl80_28 | ~ spl80_29)), inference(subsumption_resolution, [], [f8991, f607])).
fof(f8991, plain, ((e12 = e13) | (~ spl80_28 | ~ spl80_29)), inference(backward_demodulation, [], [f1421, f1417])).
fof(f1421, plain, ((e13 = op1(e13, e14)) | ~ spl80_29), inference(avatar_component_clause, [], [f1419])).
fof(f1419, plain, (spl80_29 <=> (e13 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_29])])).
fof(f8942, plain, (~ spl80_41 | ~ spl80_96), inference(avatar_contradiction_clause, [], [f8941])).
fof(f8941, plain, ($false | (~ spl80_41 | ~ spl80_96)), inference(subsumption_resolution, [], [f8940, f601])).
fof(f601, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f9])).
fof(f8940, plain, ((e10 = e12) | (~ spl80_41 | ~ spl80_96)), inference(forward_demodulation, [], [f8929, f1703])).
fof(f1703, plain, ((e10 = op1(e11, e10)) | ~ spl80_96), inference(avatar_component_clause, [], [f1701])).
fof(f1701, plain, (spl80_96 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_96])])).
fof(f8899, plain, (spl80_3 | ~ spl80_55), inference(avatar_split_clause, [], [f8896, f1528, f1310])).
fof(f1528, plain, (spl80_55 <=> (e14 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_55])])).
fof(f8896, plain, ((e12 = op1(e14, e14)) | ~ spl80_55), inference(backward_demodulation, [], [f659, f1530])).
fof(f1530, plain, ((e14 = op1(e12, e14)) | ~ spl80_55), inference(avatar_component_clause, [], [f1528])).
fof(f659, plain, (e12 = op1(op1(e12, e14), e14)), inference(cnf_transformation, [], [f12])).
fof(f8873, plain, (~ spl80_11 | ~ spl80_61), inference(avatar_split_clause, [], [f8864, f1554, f1344])).
fof(f1554, plain, (spl80_61 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_61])])).
fof(f8864, plain, (~ (e10 = op1(e14, e12)) | ~ spl80_61), inference(backward_demodulation, [], [f428, f1556])).
fof(f1556, plain, ((e10 = op1(e12, e12)) | ~ spl80_61), inference(avatar_component_clause, [], [f1554])).
fof(f428, plain, ~ (op1(e12, e12) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f8835, plain, (~ spl80_56 | ~ spl80_71), inference(avatar_split_clause, [], [f8827, f1596, f1533])).
fof(f1596, plain, (spl80_71 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_71])])).
fof(f8827, plain, (~ (e10 = op1(e12, e13)) | ~ spl80_71), inference(backward_demodulation, [], [f472, f1598])).
fof(f1598, plain, ((e10 = op1(e12, e10)) | ~ spl80_71), inference(avatar_component_clause, [], [f1596])).
fof(f8813, plain, (spl80_2 | ~ spl80_80), inference(avatar_split_clause, [], [f8809, f1633, f1306])).
fof(f1633, plain, (spl80_80 <=> (e14 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_80])])).
fof(f8809, plain, ((e11 = op1(e14, e14)) | ~ spl80_80), inference(backward_demodulation, [], [f654, f1635])).
fof(f1635, plain, ((e14 = op1(e11, e14)) | ~ spl80_80), inference(avatar_component_clause, [], [f1633])).
fof(f8797, plain, (~ spl80_56 | ~ spl80_81), inference(avatar_split_clause, [], [f8792, f1638, f1533])).
fof(f1638, plain, (spl80_81 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_81])])).
fof(f8792, plain, (~ (e10 = op1(e12, e13)) | ~ spl80_81), inference(backward_demodulation, [], [f434, f1640])).
fof(f1640, plain, ((e10 = op1(e11, e13)) | ~ spl80_81), inference(avatar_component_clause, [], [f1638])).
fof(f434, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f7])).
fof(f8784, plain, (~ spl80_11 | ~ spl80_86), inference(avatar_split_clause, [], [f8778, f1659, f1344])).
fof(f1659, plain, (spl80_86 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_86])])).
fof(f8778, plain, (~ (e10 = op1(e14, e12)) | ~ spl80_86), inference(backward_demodulation, [], [f426, f1661])).
fof(f1661, plain, ((e10 = op1(e11, e12)) | ~ spl80_86), inference(avatar_component_clause, [], [f1659])).
fof(f426, plain, ~ (op1(e11, e12) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f8589, plain, (~ spl80_752 | ~ spl80_135), inference(avatar_split_clause, [], [f8582, f1914, f5743])).
fof(f1914, plain, (spl80_135 <=> (e24 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_135])])).
fof(f8582, plain, (~ (e24 = h14(e10)) | ~ spl80_135), inference(backward_demodulation, [], [f2613, f1916])).
fof(f1916, plain, ((e24 = op2(e24, e23)) | ~ spl80_135), inference(avatar_component_clause, [], [f1914])).
fof(f2613, plain, ~ (op2(e24, e23) = h14(e10)), inference(backward_demodulation, [], [f536, f768])).
fof(f536, plain, ~ (op2(e21, e23) = op2(e24, e23)), inference(cnf_transformation, [], [f8])).
fof(f8565, plain, (~ spl80_892 | ~ spl80_140), inference(avatar_split_clause, [], [f8559, f1935, f6697])).
fof(f1935, plain, (spl80_140 <=> (e24 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_140])])).
fof(f8559, plain, (~ (e24 = h9(e10)) | ~ spl80_140), inference(backward_demodulation, [], [f2558, f1937])).
fof(f1937, plain, ((e24 = op2(e24, e22)) | ~ spl80_140), inference(avatar_component_clause, [], [f1935])).
fof(f2558, plain, ~ (op2(e24, e22) = h9(e10)), inference(backward_demodulation, [], [f523, f743])).
fof(f523, plain, ~ (op2(e20, e22) = op2(e24, e22)), inference(cnf_transformation, [], [f8])).
fof(f8555, plain, (~ spl80_142 | ~ spl80_143), inference(avatar_contradiction_clause, [], [f8554])).
fof(f8554, plain, ($false | (~ spl80_142 | ~ spl80_143)), inference(subsumption_resolution, [], [f8553, f614])).
fof(f8553, plain, ((e21 = e22) | (~ spl80_142 | ~ spl80_143)), inference(backward_demodulation, [], [f1950, f1946])).
fof(f1946, plain, ((e21 = op2(e24, e21)) | ~ spl80_142), inference(avatar_component_clause, [], [f1944])).
fof(f8545, plain, (~ spl80_974 | ~ spl80_145), inference(avatar_split_clause, [], [f8540, f1956, f7320])).
fof(f1956, plain, (spl80_145 <=> (e24 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_145])])).
fof(f8540, plain, (~ (e24 = h6(e10)) | ~ spl80_145), inference(backward_demodulation, [], [f2524, f1958])).
fof(f1958, plain, ((e24 = op2(e24, e21)) | ~ spl80_145), inference(avatar_component_clause, [], [f1956])).
fof(f2524, plain, ~ (op2(e24, e21) = h6(e10)), inference(backward_demodulation, [], [f518, f728])).
fof(f518, plain, ~ (op2(e22, e21) = op2(e24, e21)), inference(cnf_transformation, [], [f8])).
fof(f8524, plain, (spl80_1030 | ~ spl80_150), inference(avatar_split_clause, [], [f8520, f1977, f7716])).
fof(f1977, plain, (spl80_150 <=> (e24 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_150])])).
fof(f8520, plain, ((e24 = h4(e10)) | ~ spl80_150), inference(backward_demodulation, [], [f718, f1979])).
fof(f1979, plain, ((e24 = op2(e24, e20)) | ~ spl80_150), inference(avatar_component_clause, [], [f1977])).
fof(f8459, plain, (spl80_460 | ~ spl80_166), inference(avatar_split_clause, [], [f8455, f2045, f3806])).
fof(f8455, plain, ((e20 = h7(e10)) | ~ spl80_166), inference(backward_demodulation, [], [f733, f2047])).
fof(f8423, plain, (spl80_346 | ~ spl80_181), inference(avatar_split_clause, [], [f8418, f2108, f3184])).
fof(f8418, plain, ((e20 = h15(e10)) | ~ spl80_181), inference(backward_demodulation, [], [f773, f2110])).
fof(f8314, plain, (~ spl80_892 | ~ spl80_230), inference(avatar_split_clause, [], [f8310, f2313, f6697])).
fof(f2313, plain, (spl80_230 <=> (e24 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_230])])).
fof(f8310, plain, (~ (e24 = h9(e10)) | ~ spl80_230), inference(backward_demodulation, [], [f2561, f2315])).
fof(f2315, plain, ((e24 = op2(e20, e24)) | ~ spl80_230), inference(avatar_component_clause, [], [f2313])).
fof(f5136, plain, (~ spl80_641 | ~ spl80_642 | ~ spl80_643 | ~ spl80_644 | ~ spl80_645 | ~ spl80_646 | ~ spl80_647 | ~ spl80_648 | spl80_311 | spl80_307 | spl80_303 | spl80_299 | ~ spl80_650 | ~ spl80_651 | ~ spl80_652 | ~ spl80_653 | ~ spl80_654 | ~ spl80_655 | ~ spl80_656 | ~ spl80_657 | ~ spl80_658 | ~ spl80_659 | ~ spl80_660 | ~ spl80_661 | ~ spl80_662 | ~ spl80_663 | ~ spl80_664 | ~ spl80_665 | ~ spl80_666), inference(avatar_split_clause, [], [f5135, f5079, f5075, f5071, f5067, f5063, f5059, f5055, f5051, f5047, f5043, f5039, f5035, f5031, f5027, f5023, f5019, f5015, f2933, f2954, f2975, f2996, f5007, f5003, f4999, f4995, f4991, f4987, f4983, f4979])).
fof(f2996, plain, (spl80_311 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl80_311])])).
fof(f2975, plain, (spl80_307 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl80_307])])).
fof(f2954, plain, (spl80_303 <=> sP66), introduced(avatar_definition, [new_symbols(naming, [spl80_303])])).
fof(f2933, plain, (spl80_299 <=> sP67), introduced(avatar_definition, [new_symbols(naming, [spl80_299])])).
fof(f5135, plain, (~ (e20 = h17(op1(e10, e11))) | ~ (h17(e14) = h17(op1(e10, e12))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), e20)) | ~ (h17(e12) = h17(op1(e11, e10))) | ~ (op2(e24, e24) = h17(op1(e11, e11))) | ~ (h17(op1(e11, e12)) = op2(e24, h17(e12))) | ~ (h4(e10) = h17(op1(e11, e13))) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5134, f2645])).
fof(f5134, plain, (~ (h17(op1(e10, e11)) = op2(h17(e10), e24)) | ~ (h17(e14) = h17(op1(e10, e12))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), e20)) | ~ (h17(e12) = h17(op1(e11, e10))) | ~ (op2(e24, e24) = h17(op1(e11, e11))) | ~ (h17(op1(e11, e12)) = op2(e24, h17(e12))) | ~ (h4(e10) = h17(op1(e11, e13))) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5133, f781])).
fof(f5133, plain, (~ (h17(e14) = h17(op1(e10, e12))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), e20)) | ~ (h17(e12) = h17(op1(e11, e10))) | ~ (op2(e24, e24) = h17(op1(e11, e11))) | ~ (h17(op1(e11, e12)) = op2(e24, h17(e12))) | ~ (h4(e10) = h17(op1(e11, e13))) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5132, f2649])).
fof(f5132, plain, (~ (h17(op1(e10, e13)) = op2(h17(e10), e20)) | ~ (h17(e12) = h17(op1(e11, e10))) | ~ (op2(e24, e24) = h17(op1(e11, e11))) | ~ (h17(op1(e11, e12)) = op2(e24, h17(e12))) | ~ (h4(e10) = h17(op1(e11, e13))) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5131, f782])).
fof(f5131, plain, (~ (h17(e12) = h17(op1(e11, e10))) | ~ (op2(e24, e24) = h17(op1(e11, e11))) | ~ (h17(op1(e11, e12)) = op2(e24, h17(e12))) | ~ (h4(e10) = h17(op1(e11, e13))) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5130, f2646])).
fof(f5130, plain, (~ (h17(op1(e11, e10)) = op2(e24, h17(e10))) | ~ (op2(e24, e24) = h17(op1(e11, e11))) | ~ (h17(op1(e11, e12)) = op2(e24, h17(e12))) | ~ (h4(e10) = h17(op1(e11, e13))) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5129, f781])).
fof(f5129, plain, (~ (op2(e24, e24) = h17(op1(e11, e11))) | ~ (h17(op1(e11, e12)) = op2(e24, h17(e12))) | ~ (h4(e10) = h17(op1(e11, e13))) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5128, f781])).
fof(f5128, plain, (~ (h17(op1(e11, e12)) = op2(e24, h17(e12))) | ~ (h4(e10) = h17(op1(e11, e13))) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5127, f781])).
fof(f5127, plain, (~ (h4(e10) = h17(op1(e11, e13))) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5126, f718])).
fof(f5126, plain, (~ (op2(e24, e20) = h17(op1(e11, e13))) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5125, f781])).
fof(f5125, plain, (~ (h17(op1(e11, e13)) = op2(h17(e11), e20)) | ~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5124, f782])).
fof(f5124, plain, (~ (h17(op1(e11, e14)) = op2(e24, h17(e14))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5123, f781])).
fof(f5123, plain, (~ (h17(op1(e12, e11)) = op2(h17(e12), e24)) | ~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5122, f781])).
fof(f5122, plain, (~ (h17(op1(e12, e13)) = op2(h17(e12), e20)) | ~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5121, f782])).
fof(f5121, plain, (~ (h17(op1(e13, e10)) = op2(e20, h17(e10))) | ~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5120, f782])).
fof(f5120, plain, (~ (h17(e10) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5119, f783])).
fof(f5119, plain, (~ (op2(e20, e24) = h17(op1(e13, e11))) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5118, f782])).
fof(f5118, plain, (~ (h17(op1(e13, e11)) = op2(h17(e13), e24)) | ~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5117, f781])).
fof(f5117, plain, (~ (h17(op1(e13, e12)) = op2(e20, h17(e12))) | ~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5116, f782])).
fof(f5116, plain, (~ (op2(e20, e20) = h17(op1(e13, e13))) | ~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5115, f782])).
fof(f5115, plain, (~ (h17(op1(e13, e14)) = op2(e20, h17(e14))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5114, f782])).
fof(f5114, plain, (~ (h17(op1(e14, e11)) = op2(h17(e14), e24)) | ~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5113, f781])).
fof(f5113, plain, (~ (h17(op1(e14, e13)) = op2(h17(e14), e20)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(forward_demodulation, [], [f5112, f782])).
fof(f5112, plain, (sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(subsumption_resolution, [], [f1282, f781])).
fof(f1282, plain, (~ (e24 = h17(e11)) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))), inference(cnf_transformation, [], [f119])).
fof(f119, plain, (((~ (e24 = h20(e14)) & ~ (e24 = h20(e13)) & ~ (e24 = h20(e12)) & ~ (e24 = h20(e11)) & ~ (e24 = h20(e10))) | sP79 | sP78 | sP77 | sP76 | ~ (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) | ~ (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) | ~ (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) | ~ (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) | ~ (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) | ~ (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) | ~ (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) | ~ (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) | ~ (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) | ~ (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) | ~ (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) | ~ (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) | ~ (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) | ~ (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) | ~ (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) | ~ (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) | ~ (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) | ~ (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) | ~ (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) | ~ (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) | ~ (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) | ~ (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) | ~ (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) | ~ (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) | ~ (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) & ((~ (e24 = h19(e14)) & ~ (e24 = h19(e13)) & ~ (e24 = h19(e12)) & ~ (e24 = h19(e11)) & ~ (e24 = h19(e10))) | sP75 | sP74 | sP73 | sP72 | ~ (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) | ~ (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) | ~ (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) | ~ (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) | ~ (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) | ~ (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) | ~ (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) | ~ (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) | ~ (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) | ~ (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) | ~ (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) | ~ (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) | ~ (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) | ~ (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) | ~ (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) | ~ (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) | ~ (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) | ~ (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) | ~ (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) | ~ (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) | ~ (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) | ~ (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) | ~ (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) | ~ (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) | ~ (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) & ((~ (e24 = h18(e14)) & ~ (e24 = h18(e13)) & ~ (e24 = h18(e12)) & ~ (e24 = h18(e11)) & ~ (e24 = h18(e10))) | sP71 | sP70 | sP69 | sP68 | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) & ((~ (e24 = h17(e14)) & ~ (e24 = h17(e13)) & ~ (e24 = h17(e12)) & ~ (e24 = h17(e11)) & ~ (e24 = h17(e10))) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) & ((~ (e24 = h16(e14)) & ~ (e24 = h16(e13)) & ~ (e24 = h16(e12)) & ~ (e24 = h16(e11)) & ~ (e24 = h16(e10))) | sP63 | sP62 | sP61 | sP60 | ~ (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) | ~ (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) | ~ (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) | ~ (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) | ~ (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) | ~ (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) | ~ (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) | ~ (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) | ~ (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) | ~ (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) | ~ (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) | ~ (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) | ~ (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) | ~ (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) | ~ (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) | ~ (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) | ~ (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) | ~ (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) | ~ (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) | ~ (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) | ~ (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) | ~ (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) | ~ (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) | ~ (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) | ~ (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) & ((~ (e24 = h15(e14)) & ~ (e24 = h15(e13)) & ~ (e24 = h15(e12)) & ~ (e24 = h15(e11)) & ~ (e24 = h15(e10))) | sP59 | sP58 | sP57 | sP56 | ~ (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) | ~ (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) | ~ (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) | ~ (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) | ~ (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) | ~ (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) | ~ (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) | ~ (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) | ~ (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) | ~ (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) | ~ (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) | ~ (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) | ~ (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) | ~ (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) | ~ (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) | ~ (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) | ~ (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) | ~ (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) | ~ (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) | ~ (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) | ~ (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) | ~ (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) | ~ (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) | ~ (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) | ~ (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) & ((~ (e24 = h14(e14)) & ~ (e24 = h14(e13)) & ~ (e24 = h14(e12)) & ~ (e24 = h14(e11)) & ~ (e24 = h14(e10))) | sP55 | sP54 | sP53 | sP52 | ~ (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) | ~ (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) | ~ (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) | ~ (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) | ~ (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) | ~ (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) | ~ (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) | ~ (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) | ~ (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) | ~ (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) | ~ (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) | ~ (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) | ~ (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) | ~ (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) | ~ (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) | ~ (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) | ~ (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) | ~ (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) | ~ (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) | ~ (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) | ~ (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) | ~ (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) | ~ (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) | ~ (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) | ~ (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) & ((~ (e24 = h13(e14)) & ~ (e24 = h13(e13)) & ~ (e24 = h13(e12)) & ~ (e24 = h13(e11)) & ~ (e24 = h13(e10))) | sP51 | sP50 | sP49 | sP48 | ~ (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) | ~ (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) | ~ (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) | ~ (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) | ~ (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) | ~ (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) | ~ (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) | ~ (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) | ~ (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) | ~ (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) | ~ (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) | ~ (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) | ~ (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) | ~ (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) | ~ (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) | ~ (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) | ~ (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) | ~ (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) | ~ (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) | ~ (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) | ~ (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) | ~ (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) | ~ (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) | ~ (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) | ~ (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) & ((~ (e24 = h12(e14)) & ~ (e24 = h12(e13)) & ~ (e24 = h12(e12)) & ~ (e24 = h12(e11)) & ~ (e24 = h12(e10))) | sP47 | sP46 | sP45 | sP44 | ~ (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) | ~ (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) | ~ (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) | ~ (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) | ~ (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) | ~ (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e24 = h11(e14)) & ~ (e24 = h11(e13)) & ~ (e24 = h11(e12)) & ~ (e24 = h11(e11)) & ~ (e24 = h11(e10))) | sP43 | sP42 | sP41 | sP40 | ~ (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) | ~ (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) | ~ (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) | ~ (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) | ~ (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) | ~ (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e24 = h10(e14)) & ~ (e24 = h10(e13)) & ~ (e24 = h10(e12)) & ~ (e24 = h10(e11)) & ~ (e24 = h10(e10))) | sP39 | sP38 | sP37 | sP36 | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e24 = h9(e14)) & ~ (e24 = h9(e13)) & ~ (e24 = h9(e12)) & ~ (e24 = h9(e11)) & ~ (e24 = h9(e10))) | sP35 | sP34 | sP33 | sP32 | ~ (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) | ~ (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) | ~ (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) | ~ (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) | ~ (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) | ~ (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e24 = h8(e14)) & ~ (e24 = h8(e13)) & ~ (e24 = h8(e12)) & ~ (e24 = h8(e11)) & ~ (e24 = h8(e10))) | sP31 | sP30 | sP29 | sP28 | ~ (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) | ~ (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) | ~ (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) | ~ (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) | ~ (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) | ~ (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | sP23 | sP22 | sP21 | sP20 | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | sP19 | sP18 | sP17 | sP16 | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | sP15 | sP14 | sP13 | sP12 | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | sP11 | sP10 | sP9 | sP8 | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | sP7 | sP6 | sP5 | sP4 | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | sP3 | sP2 | sP1 | sP0 | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f38, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39])).
fof(f39, plain, ((~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(usedef, [], [e39])).
fof(e39, plain, (sP0 <=> (~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f40, plain, ((~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(usedef, [], [e40])).
fof(e40, plain, (sP1 <=> (~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f41, plain, ((~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(usedef, [], [e41])).
fof(e41, plain, (sP2 <=> (~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f42, plain, ((~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | ~ sP3), inference(usedef, [], [e42])).
fof(e42, plain, (sP3 <=> (~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f43, plain, ((~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP4), inference(usedef, [], [e43])).
fof(e43, plain, (sP4 <=> (~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f44, plain, ((~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP5), inference(usedef, [], [e44])).
fof(e44, plain, (sP5 <=> (~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f45, plain, ((~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP6), inference(usedef, [], [e45])).
fof(e45, plain, (sP6 <=> (~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f46, plain, ((~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP7), inference(usedef, [], [e46])).
fof(e46, plain, (sP7 <=> (~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f47, plain, ((~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP8), inference(usedef, [], [e47])).
fof(e47, plain, (sP8 <=> (~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f48, plain, ((~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP9), inference(usedef, [], [e48])).
fof(e48, plain, (sP9 <=> (~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f49, plain, ((~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP10), inference(usedef, [], [e49])).
fof(e49, plain, (sP10 <=> (~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f50, plain, ((~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP11), inference(usedef, [], [e50])).
fof(e50, plain, (sP11 <=> (~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f51, plain, ((~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP12), inference(usedef, [], [e51])).
fof(e51, plain, (sP12 <=> (~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f52, plain, ((~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP13), inference(usedef, [], [e52])).
fof(e52, plain, (sP13 <=> (~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f53, plain, ((~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP14), inference(usedef, [], [e53])).
fof(e53, plain, (sP14 <=> (~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f54, plain, ((~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | ~ sP15), inference(usedef, [], [e54])).
fof(e54, plain, (sP15 <=> (~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f55, plain, ((~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP16), inference(usedef, [], [e55])).
fof(e55, plain, (sP16 <=> (~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f56, plain, ((~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP17), inference(usedef, [], [e56])).
fof(e56, plain, (sP17 <=> (~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f57, plain, ((~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP18), inference(usedef, [], [e57])).
fof(e57, plain, (sP18 <=> (~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f58, plain, ((~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP19), inference(usedef, [], [e58])).
fof(e58, plain, (sP19 <=> (~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f59, plain, ((~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP20), inference(usedef, [], [e59])).
fof(e59, plain, (sP20 <=> (~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f60, plain, ((~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP21), inference(usedef, [], [e60])).
fof(e60, plain, (sP21 <=> (~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f61, plain, ((~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP22), inference(usedef, [], [e61])).
fof(e61, plain, (sP22 <=> (~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f62, plain, ((~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | ~ sP23), inference(usedef, [], [e62])).
fof(e62, plain, (sP23 <=> (~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f63, plain, ((~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP24), inference(usedef, [], [e63])).
fof(e63, plain, (sP24 <=> (~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f64, plain, ((~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP25), inference(usedef, [], [e64])).
fof(e64, plain, (sP25 <=> (~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f65, plain, ((~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP26), inference(usedef, [], [e65])).
fof(e65, plain, (sP26 <=> (~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f66, plain, ((~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | ~ sP27), inference(usedef, [], [e66])).
fof(e66, plain, (sP27 <=> (~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f67, plain, ((~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ sP28), inference(usedef, [], [e67])).
fof(e67, plain, (sP28 <=> (~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f68, plain, ((~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | ~ sP29), inference(usedef, [], [e68])).
fof(e68, plain, (sP29 <=> (~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f69, plain, ((~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | ~ sP30), inference(usedef, [], [e69])).
fof(e69, plain, (sP30 <=> (~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f70, plain, ((~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | ~ sP31), inference(usedef, [], [e70])).
fof(e70, plain, (sP31 <=> (~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f71, plain, ((~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ sP32), inference(usedef, [], [e71])).
fof(e71, plain, (sP32 <=> (~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f72, plain, ((~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | ~ sP33), inference(usedef, [], [e72])).
fof(e72, plain, (sP33 <=> (~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f73, plain, ((~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | ~ sP34), inference(usedef, [], [e73])).
fof(e73, plain, (sP34 <=> (~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f74, plain, ((~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | ~ sP35), inference(usedef, [], [e74])).
fof(e74, plain, (sP35 <=> (~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f75, plain, ((~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ sP36), inference(usedef, [], [e75])).
fof(e75, plain, (sP36 <=> (~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f76, plain, ((~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | ~ sP37), inference(usedef, [], [e76])).
fof(e76, plain, (sP37 <=> (~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f77, plain, ((~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | ~ sP38), inference(usedef, [], [e77])).
fof(e77, plain, (sP38 <=> (~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f78, plain, ((~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | ~ sP39), inference(usedef, [], [e78])).
fof(e78, plain, (sP39 <=> (~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f79, plain, ((~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ sP40), inference(usedef, [], [e79])).
fof(e79, plain, (sP40 <=> (~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f80, plain, ((~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | ~ sP41), inference(usedef, [], [e80])).
fof(e80, plain, (sP41 <=> (~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f81, plain, ((~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | ~ sP42), inference(usedef, [], [e81])).
fof(e81, plain, (sP42 <=> (~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f82, plain, ((~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | ~ sP43), inference(usedef, [], [e82])).
fof(e82, plain, (sP43 <=> (~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f83, plain, ((~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ sP44), inference(usedef, [], [e83])).
fof(e83, plain, (sP44 <=> (~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f84, plain, ((~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | ~ sP45), inference(usedef, [], [e84])).
fof(e84, plain, (sP45 <=> (~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f85, plain, ((~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | ~ sP46), inference(usedef, [], [e85])).
fof(e85, plain, (sP46 <=> (~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f86, plain, ((~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | ~ sP47), inference(usedef, [], [e86])).
fof(e86, plain, (sP47 <=> (~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f87, plain, ((~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10))) | ~ sP48), inference(usedef, [], [e87])).
fof(e87, plain, (sP48 <=> (~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f88, plain, ((~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10))) | ~ sP49), inference(usedef, [], [e88])).
fof(e88, plain, (sP49 <=> (~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f89, plain, ((~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10))) | ~ sP50), inference(usedef, [], [e89])).
fof(e89, plain, (sP50 <=> (~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f90, plain, ((~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10))) | ~ sP51), inference(usedef, [], [e90])).
fof(e90, plain, (sP51 <=> (~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f91, plain, ((~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10))) | ~ sP52), inference(usedef, [], [e91])).
fof(e91, plain, (sP52 <=> (~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f92, plain, ((~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10))) | ~ sP53), inference(usedef, [], [e92])).
fof(e92, plain, (sP53 <=> (~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f93, plain, ((~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10))) | ~ sP54), inference(usedef, [], [e93])).
fof(e93, plain, (sP54 <=> (~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f94, plain, ((~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10))) | ~ sP55), inference(usedef, [], [e94])).
fof(e94, plain, (sP55 <=> (~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f95, plain, ((~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10))) | ~ sP56), inference(usedef, [], [e95])).
fof(e95, plain, (sP56 <=> (~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f96, plain, ((~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10))) | ~ sP57), inference(usedef, [], [e96])).
fof(e96, plain, (sP57 <=> (~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f97, plain, ((~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10))) | ~ sP58), inference(usedef, [], [e97])).
fof(e97, plain, (sP58 <=> (~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f98, plain, ((~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10))) | ~ sP59), inference(usedef, [], [e98])).
fof(e98, plain, (sP59 <=> (~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f99, plain, ((~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10))) | ~ sP60), inference(usedef, [], [e99])).
fof(e99, plain, (sP60 <=> (~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f100, plain, ((~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10))) | ~ sP61), inference(usedef, [], [e100])).
fof(e100, plain, (sP61 <=> (~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f101, plain, ((~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10))) | ~ sP62), inference(usedef, [], [e101])).
fof(e101, plain, (sP62 <=> (~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f102, plain, ((~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10))) | ~ sP63), inference(usedef, [], [e102])).
fof(e102, plain, (sP63 <=> (~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f103, plain, ((~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10))) | ~ sP64), inference(usedef, [], [e103])).
fof(e103, plain, (sP64 <=> (~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f104, plain, ((~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10))) | ~ sP65), inference(usedef, [], [e104])).
fof(e104, plain, (sP65 <=> (~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f105, plain, ((~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10))) | ~ sP66), inference(usedef, [], [e105])).
fof(e105, plain, (sP66 <=> (~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f106, plain, ((~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10))) | ~ sP67), inference(usedef, [], [e106])).
fof(e106, plain, (sP67 <=> (~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f107, plain, ((~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10))) | ~ sP68), inference(usedef, [], [e107])).
fof(e107, plain, (sP68 <=> (~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f108, plain, ((~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10))) | ~ sP69), inference(usedef, [], [e108])).
fof(e108, plain, (sP69 <=> (~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f109, plain, ((~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10))) | ~ sP70), inference(usedef, [], [e109])).
fof(e109, plain, (sP70 <=> (~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f110, plain, ((~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10))) | ~ sP71), inference(usedef, [], [e110])).
fof(e110, plain, (sP71 <=> (~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f111, plain, ((~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10))) | ~ sP72), inference(usedef, [], [e111])).
fof(e111, plain, (sP72 <=> (~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f112, plain, ((~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10))) | ~ sP73), inference(usedef, [], [e112])).
fof(e112, plain, (sP73 <=> (~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f113, plain, ((~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10))) | ~ sP74), inference(usedef, [], [e113])).
fof(e113, plain, (sP74 <=> (~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f114, plain, ((~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10))) | ~ sP75), inference(usedef, [], [e114])).
fof(e114, plain, (sP75 <=> (~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f115, plain, ((~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10))) | ~ sP76), inference(usedef, [], [e115])).
fof(e115, plain, (sP76 <=> (~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f116, plain, ((~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10))) | ~ sP77), inference(usedef, [], [e116])).
fof(e116, plain, (sP77 <=> (~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f117, plain, ((~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10))) | ~ sP78), inference(usedef, [], [e117])).
fof(e117, plain, (sP78 <=> (~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f118, plain, ((~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10))) | ~ sP79), inference(usedef, [], [e118])).
fof(e118, plain, (sP79 <=> (~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f38, plain, (((~ (e24 = h20(e14)) & ~ (e24 = h20(e13)) & ~ (e24 = h20(e12)) & ~ (e24 = h20(e11)) & ~ (e24 = h20(e10))) | (~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10))) | (~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10))) | (~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10))) | (~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10))) | ~ (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) | ~ (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) | ~ (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) | ~ (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) | ~ (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) | ~ (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) | ~ (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) | ~ (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) | ~ (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) | ~ (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) | ~ (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) | ~ (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) | ~ (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) | ~ (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) | ~ (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) | ~ (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) | ~ (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) | ~ (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) | ~ (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) | ~ (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) | ~ (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) | ~ (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) | ~ (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) | ~ (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) | ~ (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) & ((~ (e24 = h19(e14)) & ~ (e24 = h19(e13)) & ~ (e24 = h19(e12)) & ~ (e24 = h19(e11)) & ~ (e24 = h19(e10))) | (~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10))) | (~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10))) | (~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10))) | (~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10))) | ~ (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) | ~ (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) | ~ (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) | ~ (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) | ~ (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) | ~ (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) | ~ (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) | ~ (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) | ~ (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) | ~ (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) | ~ (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) | ~ (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) | ~ (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) | ~ (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) | ~ (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) | ~ (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) | ~ (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) | ~ (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) | ~ (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) | ~ (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) | ~ (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) | ~ (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) | ~ (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) | ~ (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) | ~ (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) & ((~ (e24 = h18(e14)) & ~ (e24 = h18(e13)) & ~ (e24 = h18(e12)) & ~ (e24 = h18(e11)) & ~ (e24 = h18(e10))) | (~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10))) | (~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10))) | (~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10))) | (~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10))) | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) & ((~ (e24 = h17(e14)) & ~ (e24 = h17(e13)) & ~ (e24 = h17(e12)) & ~ (e24 = h17(e11)) & ~ (e24 = h17(e10))) | (~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10))) | (~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10))) | (~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10))) | (~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10))) | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) & ((~ (e24 = h16(e14)) & ~ (e24 = h16(e13)) & ~ (e24 = h16(e12)) & ~ (e24 = h16(e11)) & ~ (e24 = h16(e10))) | (~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10))) | (~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10))) | (~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10))) | (~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10))) | ~ (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) | ~ (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) | ~ (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) | ~ (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) | ~ (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) | ~ (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) | ~ (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) | ~ (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) | ~ (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) | ~ (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) | ~ (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) | ~ (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) | ~ (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) | ~ (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) | ~ (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) | ~ (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) | ~ (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) | ~ (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) | ~ (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) | ~ (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) | ~ (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) | ~ (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) | ~ (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) | ~ (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) | ~ (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) & ((~ (e24 = h15(e14)) & ~ (e24 = h15(e13)) & ~ (e24 = h15(e12)) & ~ (e24 = h15(e11)) & ~ (e24 = h15(e10))) | (~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10))) | (~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10))) | (~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10))) | (~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10))) | ~ (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) | ~ (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) | ~ (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) | ~ (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) | ~ (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) | ~ (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) | ~ (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) | ~ (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) | ~ (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) | ~ (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) | ~ (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) | ~ (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) | ~ (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) | ~ (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) | ~ (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) | ~ (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) | ~ (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) | ~ (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) | ~ (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) | ~ (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) | ~ (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) | ~ (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) | ~ (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) | ~ (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) | ~ (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) & ((~ (e24 = h14(e14)) & ~ (e24 = h14(e13)) & ~ (e24 = h14(e12)) & ~ (e24 = h14(e11)) & ~ (e24 = h14(e10))) | (~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10))) | (~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10))) | (~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10))) | (~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10))) | ~ (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) | ~ (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) | ~ (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) | ~ (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) | ~ (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) | ~ (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) | ~ (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) | ~ (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) | ~ (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) | ~ (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) | ~ (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) | ~ (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) | ~ (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) | ~ (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) | ~ (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) | ~ (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) | ~ (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) | ~ (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) | ~ (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) | ~ (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) | ~ (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) | ~ (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) | ~ (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) | ~ (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) | ~ (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) & ((~ (e24 = h13(e14)) & ~ (e24 = h13(e13)) & ~ (e24 = h13(e12)) & ~ (e24 = h13(e11)) & ~ (e24 = h13(e10))) | (~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10))) | (~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10))) | (~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10))) | (~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10))) | ~ (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) | ~ (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) | ~ (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) | ~ (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) | ~ (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) | ~ (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) | ~ (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) | ~ (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) | ~ (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) | ~ (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) | ~ (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) | ~ (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) | ~ (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) | ~ (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) | ~ (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) | ~ (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) | ~ (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) | ~ (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) | ~ (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) | ~ (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) | ~ (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) | ~ (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) | ~ (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) | ~ (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) | ~ (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) & ((~ (e24 = h12(e14)) & ~ (e24 = h12(e13)) & ~ (e24 = h12(e12)) & ~ (e24 = h12(e11)) & ~ (e24 = h12(e10))) | (~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | (~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | (~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | (~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) | ~ (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) | ~ (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) | ~ (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) | ~ (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) | ~ (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e24 = h11(e14)) & ~ (e24 = h11(e13)) & ~ (e24 = h11(e12)) & ~ (e24 = h11(e11)) & ~ (e24 = h11(e10))) | (~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | (~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | (~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | (~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) | ~ (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) | ~ (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) | ~ (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) | ~ (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) | ~ (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e24 = h10(e14)) & ~ (e24 = h10(e13)) & ~ (e24 = h10(e12)) & ~ (e24 = h10(e11)) & ~ (e24 = h10(e10))) | (~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | (~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | (~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | (~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e24 = h9(e14)) & ~ (e24 = h9(e13)) & ~ (e24 = h9(e12)) & ~ (e24 = h9(e11)) & ~ (e24 = h9(e10))) | (~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | (~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | (~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | (~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) | ~ (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) | ~ (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) | ~ (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) | ~ (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) | ~ (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e24 = h8(e14)) & ~ (e24 = h8(e13)) & ~ (e24 = h8(e12)) & ~ (e24 = h8(e11)) & ~ (e24 = h8(e10))) | (~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | (~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | (~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | (~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) | ~ (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) | ~ (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) | ~ (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) | ~ (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) | ~ (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | (~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | (~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | (~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | (~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | (~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | (~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | (~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | (~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | (~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | (~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ~ ((((e24 = h20(e14)) | (e24 = h20(e13)) | (e24 = h20(e12)) | (e24 = h20(e11)) | (e24 = h20(e10))) & ((e23 = h20(e14)) | (e23 = h20(e13)) | (e23 = h20(e12)) | (e23 = h20(e11)) | (e23 = h20(e10))) & ((e22 = h20(e14)) | (e22 = h20(e13)) | (e22 = h20(e12)) | (e22 = h20(e11)) | (e22 = h20(e10))) & ((e21 = h20(e14)) | (e21 = h20(e13)) | (e21 = h20(e12)) | (e21 = h20(e11)) | (e21 = h20(e10))) & ((e20 = h20(e14)) | (e20 = h20(e13)) | (e20 = h20(e12)) | (e20 = h20(e11)) | (e20 = h20(e10))) & (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) & (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) & (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) & (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) & (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) & (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) & (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) & (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) & (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) & (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) & (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) & (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) & (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) & (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) & (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) & (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) & (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) & (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) & (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) & (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) & (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) & (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) & (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) & (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) & (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) | (((e24 = h19(e14)) | (e24 = h19(e13)) | (e24 = h19(e12)) | (e24 = h19(e11)) | (e24 = h19(e10))) & ((e23 = h19(e14)) | (e23 = h19(e13)) | (e23 = h19(e12)) | (e23 = h19(e11)) | (e23 = h19(e10))) & ((e22 = h19(e14)) | (e22 = h19(e13)) | (e22 = h19(e12)) | (e22 = h19(e11)) | (e22 = h19(e10))) & ((e21 = h19(e14)) | (e21 = h19(e13)) | (e21 = h19(e12)) | (e21 = h19(e11)) | (e21 = h19(e10))) & ((e20 = h19(e14)) | (e20 = h19(e13)) | (e20 = h19(e12)) | (e20 = h19(e11)) | (e20 = h19(e10))) & (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) & (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) & (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) & (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) & (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) & (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) & (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) & (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) & (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) & (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) & (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) & (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) & (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) & (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) & (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) & (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) & (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) & (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) & (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) & (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) & (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) & (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) & (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) & (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) & (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) | (((e24 = h18(e14)) | (e24 = h18(e13)) | (e24 = h18(e12)) | (e24 = h18(e11)) | (e24 = h18(e10))) & ((e23 = h18(e14)) | (e23 = h18(e13)) | (e23 = h18(e12)) | (e23 = h18(e11)) | (e23 = h18(e10))) & ((e22 = h18(e14)) | (e22 = h18(e13)) | (e22 = h18(e12)) | (e22 = h18(e11)) | (e22 = h18(e10))) & ((e21 = h18(e14)) | (e21 = h18(e13)) | (e21 = h18(e12)) | (e21 = h18(e11)) | (e21 = h18(e10))) & ((e20 = h18(e14)) | (e20 = h18(e13)) | (e20 = h18(e12)) | (e20 = h18(e11)) | (e20 = h18(e10))) & (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) & (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) & (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) & (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) & (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) & (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) & (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) & (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) & (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) & (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) & (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) & (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) & (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) & (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) & (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) & (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) & (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) & (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) & (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) & (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) & (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) & (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) & (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) & (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) & (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) | (((e24 = h17(e14)) | (e24 = h17(e13)) | (e24 = h17(e12)) | (e24 = h17(e11)) | (e24 = h17(e10))) & ((e23 = h17(e14)) | (e23 = h17(e13)) | (e23 = h17(e12)) | (e23 = h17(e11)) | (e23 = h17(e10))) & ((e22 = h17(e14)) | (e22 = h17(e13)) | (e22 = h17(e12)) | (e22 = h17(e11)) | (e22 = h17(e10))) & ((e21 = h17(e14)) | (e21 = h17(e13)) | (e21 = h17(e12)) | (e21 = h17(e11)) | (e21 = h17(e10))) & ((e20 = h17(e14)) | (e20 = h17(e13)) | (e20 = h17(e12)) | (e20 = h17(e11)) | (e20 = h17(e10))) & (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) & (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) & (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) & (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) & (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) & (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) & (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) & (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) & (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) & (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) & (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) & (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) & (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) & (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) & (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) & (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) & (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) & (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) & (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) & (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) & (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) & (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) & (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) & (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) & (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) | (((e24 = h16(e14)) | (e24 = h16(e13)) | (e24 = h16(e12)) | (e24 = h16(e11)) | (e24 = h16(e10))) & ((e23 = h16(e14)) | (e23 = h16(e13)) | (e23 = h16(e12)) | (e23 = h16(e11)) | (e23 = h16(e10))) & ((e22 = h16(e14)) | (e22 = h16(e13)) | (e22 = h16(e12)) | (e22 = h16(e11)) | (e22 = h16(e10))) & ((e21 = h16(e14)) | (e21 = h16(e13)) | (e21 = h16(e12)) | (e21 = h16(e11)) | (e21 = h16(e10))) & ((e20 = h16(e14)) | (e20 = h16(e13)) | (e20 = h16(e12)) | (e20 = h16(e11)) | (e20 = h16(e10))) & (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) & (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) & (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) & (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) & (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) & (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) & (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) & (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) & (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) & (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) & (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) & (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) & (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) & (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) & (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) & (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) & (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) & (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) & (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) & (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) & (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) & (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) & (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) & (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) & (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) | (((e24 = h15(e14)) | (e24 = h15(e13)) | (e24 = h15(e12)) | (e24 = h15(e11)) | (e24 = h15(e10))) & ((e23 = h15(e14)) | (e23 = h15(e13)) | (e23 = h15(e12)) | (e23 = h15(e11)) | (e23 = h15(e10))) & ((e22 = h15(e14)) | (e22 = h15(e13)) | (e22 = h15(e12)) | (e22 = h15(e11)) | (e22 = h15(e10))) & ((e21 = h15(e14)) | (e21 = h15(e13)) | (e21 = h15(e12)) | (e21 = h15(e11)) | (e21 = h15(e10))) & ((e20 = h15(e14)) | (e20 = h15(e13)) | (e20 = h15(e12)) | (e20 = h15(e11)) | (e20 = h15(e10))) & (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) & (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) & (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) & (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) & (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) & (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) & (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) & (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) & (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) & (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) & (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) & (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) & (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) & (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) & (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) & (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) & (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) & (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) & (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) & (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) & (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) & (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) & (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) & (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) & (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) | (((e24 = h14(e14)) | (e24 = h14(e13)) | (e24 = h14(e12)) | (e24 = h14(e11)) | (e24 = h14(e10))) & ((e23 = h14(e14)) | (e23 = h14(e13)) | (e23 = h14(e12)) | (e23 = h14(e11)) | (e23 = h14(e10))) & ((e22 = h14(e14)) | (e22 = h14(e13)) | (e22 = h14(e12)) | (e22 = h14(e11)) | (e22 = h14(e10))) & ((e21 = h14(e14)) | (e21 = h14(e13)) | (e21 = h14(e12)) | (e21 = h14(e11)) | (e21 = h14(e10))) & ((e20 = h14(e14)) | (e20 = h14(e13)) | (e20 = h14(e12)) | (e20 = h14(e11)) | (e20 = h14(e10))) & (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) & (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) & (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) & (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) & (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) & (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) & (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) & (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) & (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) & (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) & (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) & (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) & (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) & (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) & (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) & (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) & (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) & (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) & (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) & (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) & (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) & (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) & (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) & (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) & (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) | (((e24 = h13(e14)) | (e24 = h13(e13)) | (e24 = h13(e12)) | (e24 = h13(e11)) | (e24 = h13(e10))) & ((e23 = h13(e14)) | (e23 = h13(e13)) | (e23 = h13(e12)) | (e23 = h13(e11)) | (e23 = h13(e10))) & ((e22 = h13(e14)) | (e22 = h13(e13)) | (e22 = h13(e12)) | (e22 = h13(e11)) | (e22 = h13(e10))) & ((e21 = h13(e14)) | (e21 = h13(e13)) | (e21 = h13(e12)) | (e21 = h13(e11)) | (e21 = h13(e10))) & ((e20 = h13(e14)) | (e20 = h13(e13)) | (e20 = h13(e12)) | (e20 = h13(e11)) | (e20 = h13(e10))) & (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) & (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) & (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) & (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) & (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) & (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) & (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) & (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) & (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) & (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) & (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) & (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) & (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) & (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) & (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) & (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) & (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) & (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) & (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) & (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) & (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) & (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) & (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) & (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) & (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) | (((e24 = h12(e14)) | (e24 = h12(e13)) | (e24 = h12(e12)) | (e24 = h12(e11)) | (e24 = h12(e10))) & ((e23 = h12(e14)) | (e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e14)) | (e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e14)) | (e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e14)) | (e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) & (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) & (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) & (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) & (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) & (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e24 = h11(e14)) | (e24 = h11(e13)) | (e24 = h11(e12)) | (e24 = h11(e11)) | (e24 = h11(e10))) & ((e23 = h11(e14)) | (e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e14)) | (e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e14)) | (e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e14)) | (e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) & (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) & (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) & (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) & (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) & (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e24 = h10(e14)) | (e24 = h10(e13)) | (e24 = h10(e12)) | (e24 = h10(e11)) | (e24 = h10(e10))) & ((e23 = h10(e14)) | (e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e14)) | (e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e14)) | (e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e14)) | (e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) & (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) & (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) & (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) & (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) & (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e24 = h9(e14)) | (e24 = h9(e13)) | (e24 = h9(e12)) | (e24 = h9(e11)) | (e24 = h9(e10))) & ((e23 = h9(e14)) | (e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e14)) | (e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e14)) | (e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e14)) | (e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) & (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) & (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) & (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) & (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) & (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e24 = h8(e14)) | (e24 = h8(e13)) | (e24 = h8(e12)) | (e24 = h8(e11)) | (e24 = h8(e10))) & ((e23 = h8(e14)) | (e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e14)) | (e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e14)) | (e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e14)) | (e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) & (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) & (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) & (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) & (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) & (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f36])).
fof(f36, plain, ~ ((((e24 = h20(e14)) | (e24 = h20(e13)) | (e24 = h20(e12)) | (e24 = h20(e11)) | (e24 = h20(e10))) & ((e23 = h20(e14)) | (e23 = h20(e13)) | (e23 = h20(e12)) | (e23 = h20(e11)) | (e23 = h20(e10))) & ((e22 = h20(e14)) | (e22 = h20(e13)) | (e22 = h20(e12)) | (e22 = h20(e11)) | (e22 = h20(e10))) & ((e21 = h20(e14)) | (e21 = h20(e13)) | (e21 = h20(e12)) | (e21 = h20(e11)) | (e21 = h20(e10))) & ((e20 = h20(e14)) | (e20 = h20(e13)) | (e20 = h20(e12)) | (e20 = h20(e11)) | (e20 = h20(e10))) & (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) & (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) & (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) & (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) & (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) & (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) & (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) & (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) & (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) & (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) & (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) & (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) & (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) & (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) & (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) & (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) & (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) & (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) & (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) & (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) & (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) & (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) & (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) & (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) & (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) | (((e24 = h19(e14)) | (e24 = h19(e13)) | (e24 = h19(e12)) | (e24 = h19(e11)) | (e24 = h19(e10))) & ((e23 = h19(e14)) | (e23 = h19(e13)) | (e23 = h19(e12)) | (e23 = h19(e11)) | (e23 = h19(e10))) & ((e22 = h19(e14)) | (e22 = h19(e13)) | (e22 = h19(e12)) | (e22 = h19(e11)) | (e22 = h19(e10))) & ((e21 = h19(e14)) | (e21 = h19(e13)) | (e21 = h19(e12)) | (e21 = h19(e11)) | (e21 = h19(e10))) & ((e20 = h19(e14)) | (e20 = h19(e13)) | (e20 = h19(e12)) | (e20 = h19(e11)) | (e20 = h19(e10))) & (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) & (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) & (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) & (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) & (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) & (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) & (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) & (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) & (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) & (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) & (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) & (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) & (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) & (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) & (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) & (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) & (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) & (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) & (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) & (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) & (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) & (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) & (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) & (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) & (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) | (((e24 = h18(e14)) | (e24 = h18(e13)) | (e24 = h18(e12)) | (e24 = h18(e11)) | (e24 = h18(e10))) & ((e23 = h18(e14)) | (e23 = h18(e13)) | (e23 = h18(e12)) | (e23 = h18(e11)) | (e23 = h18(e10))) & ((e22 = h18(e14)) | (e22 = h18(e13)) | (e22 = h18(e12)) | (e22 = h18(e11)) | (e22 = h18(e10))) & ((e21 = h18(e14)) | (e21 = h18(e13)) | (e21 = h18(e12)) | (e21 = h18(e11)) | (e21 = h18(e10))) & ((e20 = h18(e14)) | (e20 = h18(e13)) | (e20 = h18(e12)) | (e20 = h18(e11)) | (e20 = h18(e10))) & (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) & (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) & (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) & (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) & (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) & (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) & (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) & (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) & (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) & (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) & (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) & (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) & (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) & (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) & (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) & (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) & (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) & (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) & (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) & (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) & (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) & (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) & (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) & (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) & (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) | (((e24 = h17(e14)) | (e24 = h17(e13)) | (e24 = h17(e12)) | (e24 = h17(e11)) | (e24 = h17(e10))) & ((e23 = h17(e14)) | (e23 = h17(e13)) | (e23 = h17(e12)) | (e23 = h17(e11)) | (e23 = h17(e10))) & ((e22 = h17(e14)) | (e22 = h17(e13)) | (e22 = h17(e12)) | (e22 = h17(e11)) | (e22 = h17(e10))) & ((e21 = h17(e14)) | (e21 = h17(e13)) | (e21 = h17(e12)) | (e21 = h17(e11)) | (e21 = h17(e10))) & ((e20 = h17(e14)) | (e20 = h17(e13)) | (e20 = h17(e12)) | (e20 = h17(e11)) | (e20 = h17(e10))) & (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) & (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) & (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) & (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) & (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) & (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) & (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) & (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) & (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) & (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) & (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) & (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) & (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) & (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) & (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) & (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) & (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) & (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) & (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) & (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) & (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) & (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) & (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) & (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) & (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) | (((e24 = h16(e14)) | (e24 = h16(e13)) | (e24 = h16(e12)) | (e24 = h16(e11)) | (e24 = h16(e10))) & ((e23 = h16(e14)) | (e23 = h16(e13)) | (e23 = h16(e12)) | (e23 = h16(e11)) | (e23 = h16(e10))) & ((e22 = h16(e14)) | (e22 = h16(e13)) | (e22 = h16(e12)) | (e22 = h16(e11)) | (e22 = h16(e10))) & ((e21 = h16(e14)) | (e21 = h16(e13)) | (e21 = h16(e12)) | (e21 = h16(e11)) | (e21 = h16(e10))) & ((e20 = h16(e14)) | (e20 = h16(e13)) | (e20 = h16(e12)) | (e20 = h16(e11)) | (e20 = h16(e10))) & (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) & (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) & (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) & (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) & (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) & (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) & (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) & (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) & (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) & (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) & (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) & (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) & (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) & (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) & (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) & (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) & (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) & (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) & (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) & (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) & (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) & (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) & (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) & (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) & (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) | (((e24 = h15(e14)) | (e24 = h15(e13)) | (e24 = h15(e12)) | (e24 = h15(e11)) | (e24 = h15(e10))) & ((e23 = h15(e14)) | (e23 = h15(e13)) | (e23 = h15(e12)) | (e23 = h15(e11)) | (e23 = h15(e10))) & ((e22 = h15(e14)) | (e22 = h15(e13)) | (e22 = h15(e12)) | (e22 = h15(e11)) | (e22 = h15(e10))) & ((e21 = h15(e14)) | (e21 = h15(e13)) | (e21 = h15(e12)) | (e21 = h15(e11)) | (e21 = h15(e10))) & ((e20 = h15(e14)) | (e20 = h15(e13)) | (e20 = h15(e12)) | (e20 = h15(e11)) | (e20 = h15(e10))) & (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) & (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) & (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) & (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) & (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) & (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) & (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) & (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) & (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) & (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) & (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) & (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) & (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) & (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) & (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) & (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) & (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) & (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) & (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) & (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) & (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) & (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) & (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) & (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) & (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) | (((e24 = h14(e14)) | (e24 = h14(e13)) | (e24 = h14(e12)) | (e24 = h14(e11)) | (e24 = h14(e10))) & ((e23 = h14(e14)) | (e23 = h14(e13)) | (e23 = h14(e12)) | (e23 = h14(e11)) | (e23 = h14(e10))) & ((e22 = h14(e14)) | (e22 = h14(e13)) | (e22 = h14(e12)) | (e22 = h14(e11)) | (e22 = h14(e10))) & ((e21 = h14(e14)) | (e21 = h14(e13)) | (e21 = h14(e12)) | (e21 = h14(e11)) | (e21 = h14(e10))) & ((e20 = h14(e14)) | (e20 = h14(e13)) | (e20 = h14(e12)) | (e20 = h14(e11)) | (e20 = h14(e10))) & (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) & (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) & (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) & (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) & (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) & (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) & (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) & (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) & (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) & (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) & (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) & (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) & (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) & (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) & (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) & (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) & (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) & (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) & (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) & (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) & (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) & (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) & (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) & (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) & (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) | (((e24 = h13(e14)) | (e24 = h13(e13)) | (e24 = h13(e12)) | (e24 = h13(e11)) | (e24 = h13(e10))) & ((e23 = h13(e14)) | (e23 = h13(e13)) | (e23 = h13(e12)) | (e23 = h13(e11)) | (e23 = h13(e10))) & ((e22 = h13(e14)) | (e22 = h13(e13)) | (e22 = h13(e12)) | (e22 = h13(e11)) | (e22 = h13(e10))) & ((e21 = h13(e14)) | (e21 = h13(e13)) | (e21 = h13(e12)) | (e21 = h13(e11)) | (e21 = h13(e10))) & ((e20 = h13(e14)) | (e20 = h13(e13)) | (e20 = h13(e12)) | (e20 = h13(e11)) | (e20 = h13(e10))) & (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) & (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) & (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) & (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) & (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) & (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) & (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) & (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) & (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) & (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) & (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) & (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) & (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) & (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) & (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) & (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) & (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) & (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) & (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) & (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) & (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) & (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) & (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) & (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) & (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) | (((e24 = h12(e14)) | (e24 = h12(e13)) | (e24 = h12(e12)) | (e24 = h12(e11)) | (e24 = h12(e10))) & ((e23 = h12(e14)) | (e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e14)) | (e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e14)) | (e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e14)) | (e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) & (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) & (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) & (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) & (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) & (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e24 = h11(e14)) | (e24 = h11(e13)) | (e24 = h11(e12)) | (e24 = h11(e11)) | (e24 = h11(e10))) & ((e23 = h11(e14)) | (e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e14)) | (e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e14)) | (e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e14)) | (e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) & (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) & (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) & (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) & (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) & (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e24 = h10(e14)) | (e24 = h10(e13)) | (e24 = h10(e12)) | (e24 = h10(e11)) | (e24 = h10(e10))) & ((e23 = h10(e14)) | (e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e14)) | (e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e14)) | (e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e14)) | (e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) & (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) & (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) & (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) & (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) & (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e24 = h9(e14)) | (e24 = h9(e13)) | (e24 = h9(e12)) | (e24 = h9(e11)) | (e24 = h9(e10))) & ((e23 = h9(e14)) | (e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e14)) | (e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e14)) | (e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e14)) | (e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) & (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) & (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) & (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) & (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) & (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e24 = h8(e14)) | (e24 = h8(e13)) | (e24 = h8(e12)) | (e24 = h8(e11)) | (e24 = h8(e10))) & ((e23 = h8(e14)) | (e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e14)) | (e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e14)) | (e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e14)) | (e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) & (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) & (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) & (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) & (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) & (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', co1)).
fof(f3005, plain, ~ spl80_311, inference(avatar_split_clause, [], [f3004, f2996])).
fof(f3004, plain, ~ sP64, inference(subsumption_resolution, [], [f879, f782])).
fof(f879, plain, (~ (e20 = h17(e13)) | ~ sP64), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10))) | ~ sP64), inference(nnf_transformation, [], [f103])).
fof(f2994, plain, (~ spl80_307 | ~ spl80_310), inference(avatar_split_clause, [], [f871, f2991, f2975])).
fof(f871, plain, (~ (e21 = h17(e10)) | ~ sP65), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10))) | ~ sP65), inference(nnf_transformation, [], [f104])).
fof(f2967, plain, (~ spl80_303 | ~ spl80_305), inference(avatar_split_clause, [], [f868, f2964, f2954])).
fof(f868, plain, (~ (e22 = h17(e12)) | ~ sP66), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10))) | ~ sP66), inference(nnf_transformation, [], [f105])).
fof(f2940, plain, (~ spl80_299 | ~ spl80_300), inference(avatar_split_clause, [], [f865, f2937, f2933])).
fof(f865, plain, (~ (e23 = h17(e14)) | ~ sP67), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10))) | ~ sP67), inference(nnf_transformation, [], [f106])).
fof(f2454, plain, spl80_166, inference(avatar_split_clause, [], [f698, f2045])).
fof(f698, plain, (e20 = op2(e23, e21)), inference(cnf_transformation, [], [f15])).
fof(f2452, plain, spl80_41, inference(avatar_split_clause, [], [f695, f1470])).
fof(f695, plain, (e10 = op1(e13, e11)), inference(cnf_transformation, [], [f14])).
fof(f2441, plain, (spl80_250 | spl80_225 | spl80_200 | spl80_175 | spl80_150), inference(avatar_split_clause, [], [f334, f1977, f2082, f2187, f2292, f2397])).
fof(f334, plain, ((e24 = op2(e24, e20)) | (e24 = op2(e23, e20)) | (e24 = op2(e22, e20)) | (e24 = op2(e21, e20)) | (op2(e20, e20) = e24)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e24 = op2(e24, e24)) | (e24 = op2(e23, e24)) | (e24 = op2(e22, e24)) | (e24 = op2(e21, e24)) | (e24 = op2(e20, e24))) & ((e24 = op2(e24, e24)) | (e24 = op2(e24, e23)) | (e24 = op2(e24, e22)) | (e24 = op2(e24, e21)) | (e24 = op2(e24, e20))) & ((e23 = op2(e24, e24)) | (e23 = op2(e23, e24)) | (e23 = op2(e22, e24)) | (e23 = op2(e21, e24)) | (e23 = op2(e20, e24))) & ((e23 = op2(e24, e24)) | (e23 = op2(e24, e23)) | (e23 = op2(e24, e22)) | (e23 = op2(e24, e21)) | (e23 = op2(e24, e20))) & ((e22 = op2(e24, e24)) | (e22 = op2(e23, e24)) | (e22 = op2(e22, e24)) | (e22 = op2(e21, e24)) | (e22 = op2(e20, e24))) & ((e22 = op2(e24, e24)) | (e22 = op2(e24, e23)) | (e22 = op2(e24, e22)) | (e22 = op2(e24, e21)) | (e22 = op2(e24, e20))) & ((e21 = op2(e24, e24)) | (e21 = op2(e23, e24)) | (e21 = op2(e22, e24)) | (e21 = op2(e21, e24)) | (e21 = op2(e20, e24))) & ((e21 = op2(e24, e24)) | (e21 = op2(e24, e23)) | (e21 = op2(e24, e22)) | (e21 = op2(e24, e21)) | (e21 = op2(e24, e20))) & ((e20 = op2(e24, e24)) | (e20 = op2(e23, e24)) | (e20 = op2(e22, e24)) | (e20 = op2(e21, e24)) | (e20 = op2(e20, e24))) & ((e20 = op2(e24, e24)) | (e20 = op2(e24, e23)) | (e20 = op2(e24, e22)) | (e20 = op2(e24, e21)) | (e20 = op2(e24, e20))) & ((e24 = op2(e24, e23)) | (e24 = op2(e23, e23)) | (e24 = op2(e22, e23)) | (e24 = op2(e21, e23)) | (e24 = op2(e20, e23))) & ((e24 = op2(e23, e24)) | (e24 = op2(e23, e23)) | (e24 = op2(e23, e22)) | (e24 = op2(e23, e21)) | (e24 = op2(e23, e20))) & ((e23 = op2(e24, e23)) | (e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e24)) | (e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e24, e23)) | (e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e24)) | (e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e24, e23)) | (e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e24)) | (e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e24, e23)) | (e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e24)) | (e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e24 = op2(e24, e22)) | (e24 = op2(e23, e22)) | (e24 = op2(e22, e22)) | (e24 = op2(e21, e22)) | (e24 = op2(e20, e22))) & ((e24 = op2(e22, e24)) | (e24 = op2(e22, e23)) | (e24 = op2(e22, e22)) | (e24 = op2(e22, e21)) | (e24 = op2(e22, e20))) & ((e23 = op2(e24, e22)) | (e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e24)) | (e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e24, e22)) | (e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e24)) | (e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e24, e22)) | (e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e24)) | (e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e24, e22)) | (e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e24 = op2(e24, e21)) | (e24 = op2(e23, e21)) | (e24 = op2(e22, e21)) | (e24 = op2(e21, e21)) | (e24 = op2(e20, e21))) & ((e24 = op2(e21, e24)) | (e24 = op2(e21, e23)) | (e24 = op2(e21, e22)) | (e24 = op2(e21, e21)) | (e24 = op2(e21, e20))) & ((e23 = op2(e24, e21)) | (e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e24)) | (e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e24, e21)) | (e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e24)) | (e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e24, e21)) | (e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e24)) | (e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e24, e21)) | (e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e24)) | (e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e24 = op2(e24, e20)) | (e24 = op2(e23, e20)) | (e24 = op2(e22, e20)) | (e24 = op2(e21, e20)) | (op2(e20, e20) = e24)) & ((e24 = op2(e20, e24)) | (e24 = op2(e20, e23)) | (e24 = op2(e20, e22)) | (e24 = op2(e20, e21)) | (op2(e20, e20) = e24)) & ((e23 = op2(e24, e20)) | (e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e24)) | (e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e24, e20)) | (e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e24)) | (e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e24, e20)) | (e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e24)) | (e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e24, e20)) | (e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e24)) | (e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax5)).
fof(f2427, plain, (spl80_237 | spl80_212 | spl80_187 | spl80_162 | spl80_137), inference(avatar_split_clause, [], [f348, f1923, f2028, f2133, f2238, f2343])).
fof(f348, plain, ((e21 = op2(e24, e22)) | (e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))), inference(cnf_transformation, [], [f5])).
fof(f2426, plain, (spl80_198 | spl80_193 | spl80_188 | spl80_183 | spl80_178), inference(avatar_split_clause, [], [f349, f2095, f2116, f2137, f2158, f2179])).
fof(f349, plain, ((e22 = op2(e22, e24)) | (e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))), inference(cnf_transformation, [], [f5])).
fof(f2416, plain, (spl80_173 | spl80_168 | spl80_163 | spl80_158 | spl80_153), inference(avatar_split_clause, [], [f359, f1990, f2011, f2032, f2053, f2074])).
fof(f359, plain, ((e22 = op2(e23, e24)) | (e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))), inference(cnf_transformation, [], [f5])).
fof(f2413, plain, (spl80_234 | spl80_209 | spl80_184 | spl80_159 | spl80_134), inference(avatar_split_clause, [], [f362, f1910, f2015, f2120, f2225, f2330])).
fof(f362, plain, ((e23 = op2(e24, e23)) | (e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))), inference(cnf_transformation, [], [f5])).
fof(f2408, plain, (spl80_147 | spl80_142 | spl80_137 | spl80_132 | spl80_127), inference(avatar_split_clause, [], [f367, f1881, f1902, f1923, f1944, f1965])).
fof(f367, plain, ((e21 = op2(e24, e24)) | (e21 = op2(e24, e23)) | (e21 = op2(e24, e22)) | (e21 = op2(e24, e21)) | (e21 = op2(e24, e20))), inference(cnf_transformation, [], [f5])).
fof(f2402, plain, (spl80_150 | spl80_145 | spl80_140 | spl80_135 | spl80_130), inference(avatar_split_clause, [], [f373, f1893, f1914, f1935, f1956, f1977])).
fof(f373, plain, ((e24 = op2(e24, e24)) | (e24 = op2(e24, e23)) | (e24 = op2(e24, e22)) | (e24 = op2(e24, e21)) | (e24 = op2(e24, e20))), inference(cnf_transformation, [], [f5])).
fof(f2400, plain, (spl80_246 | spl80_247 | spl80_248 | spl80_249 | spl80_250), inference(avatar_split_clause, [], [f300, f2397, f2393, f2389, f2385, f2381])).
fof(f300, plain, ((op2(e20, e20) = e24) | (op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e24 = op2(e24, e24)) | (e23 = op2(e24, e24)) | (e22 = op2(e24, e24)) | (e21 = op2(e24, e24)) | (e20 = op2(e24, e24))) & ((e24 = op2(e24, e23)) | (e23 = op2(e24, e23)) | (e22 = op2(e24, e23)) | (e21 = op2(e24, e23)) | (e20 = op2(e24, e23))) & ((e24 = op2(e24, e22)) | (e23 = op2(e24, e22)) | (e22 = op2(e24, e22)) | (e21 = op2(e24, e22)) | (e20 = op2(e24, e22))) & ((e24 = op2(e24, e21)) | (e23 = op2(e24, e21)) | (e22 = op2(e24, e21)) | (e21 = op2(e24, e21)) | (e20 = op2(e24, e21))) & ((e24 = op2(e24, e20)) | (e23 = op2(e24, e20)) | (e22 = op2(e24, e20)) | (e21 = op2(e24, e20)) | (e20 = op2(e24, e20))) & ((e24 = op2(e23, e24)) | (e23 = op2(e23, e24)) | (e22 = op2(e23, e24)) | (e21 = op2(e23, e24)) | (e20 = op2(e23, e24))) & ((e24 = op2(e23, e23)) | (e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e24 = op2(e23, e22)) | (e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e24 = op2(e23, e21)) | (e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e24 = op2(e23, e20)) | (e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e24 = op2(e22, e24)) | (e23 = op2(e22, e24)) | (e22 = op2(e22, e24)) | (e21 = op2(e22, e24)) | (e20 = op2(e22, e24))) & ((e24 = op2(e22, e23)) | (e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e24 = op2(e22, e22)) | (e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e24 = op2(e22, e21)) | (e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e24 = op2(e22, e20)) | (e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e24 = op2(e21, e24)) | (e23 = op2(e21, e24)) | (e22 = op2(e21, e24)) | (e21 = op2(e21, e24)) | (e20 = op2(e21, e24))) & ((e24 = op2(e21, e23)) | (e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e24 = op2(e21, e22)) | (e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e24 = op2(e21, e21)) | (e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e24 = op2(e21, e20)) | (e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e24 = op2(e20, e24)) | (e23 = op2(e20, e24)) | (e22 = op2(e20, e24)) | (e21 = op2(e20, e24)) | (e20 = op2(e20, e24))) & ((e24 = op2(e20, e23)) | (e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e24 = op2(e20, e22)) | (e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e24 = op2(e20, e21)) | (e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e24) | (op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax4)).
fof(f2316, plain, (spl80_226 | spl80_227 | spl80_228 | spl80_229 | spl80_230), inference(avatar_split_clause, [], [f304, f2313, f2309, f2305, f2301, f2297])).
fof(f304, plain, ((e24 = op2(e20, e24)) | (e23 = op2(e20, e24)) | (e22 = op2(e20, e24)) | (e21 = op2(e20, e24)) | (e20 = op2(e20, e24))), inference(cnf_transformation, [], [f4])).
fof(f2274, plain, (spl80_216 | spl80_217 | spl80_218 | spl80_219 | spl80_220), inference(avatar_split_clause, [], [f306, f2271, f2267, f2263, f2259, f2255])).
fof(f306, plain, ((e24 = op2(e21, e21)) | (e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))), inference(cnf_transformation, [], [f4])).
fof(f1874, plain, (spl80_121 | spl80_96 | spl80_71 | spl80_46 | spl80_21), inference(avatar_split_clause, [], [f226, f1386, f1491, f1596, f1701, f1806])).
fof(f226, plain, ((e10 = op1(e14, e10)) | (e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e14 = op1(e14, e14)) | (e14 = op1(e13, e14)) | (e14 = op1(e12, e14)) | (e14 = op1(e11, e14)) | (e14 = op1(e10, e14))) & ((e14 = op1(e14, e14)) | (e14 = op1(e14, e13)) | (e14 = op1(e14, e12)) | (e14 = op1(e14, e11)) | (e14 = op1(e14, e10))) & ((e13 = op1(e14, e14)) | (e13 = op1(e13, e14)) | (e13 = op1(e12, e14)) | (e13 = op1(e11, e14)) | (e13 = op1(e10, e14))) & ((e13 = op1(e14, e14)) | (e13 = op1(e14, e13)) | (e13 = op1(e14, e12)) | (e13 = op1(e14, e11)) | (e13 = op1(e14, e10))) & ((e12 = op1(e14, e14)) | (e12 = op1(e13, e14)) | (e12 = op1(e12, e14)) | (e12 = op1(e11, e14)) | (e12 = op1(e10, e14))) & ((e12 = op1(e14, e14)) | (e12 = op1(e14, e13)) | (e12 = op1(e14, e12)) | (e12 = op1(e14, e11)) | (e12 = op1(e14, e10))) & ((e11 = op1(e14, e14)) | (e11 = op1(e13, e14)) | (e11 = op1(e12, e14)) | (e11 = op1(e11, e14)) | (e11 = op1(e10, e14))) & ((e11 = op1(e14, e14)) | (e11 = op1(e14, e13)) | (e11 = op1(e14, e12)) | (e11 = op1(e14, e11)) | (e11 = op1(e14, e10))) & ((e10 = op1(e14, e14)) | (e10 = op1(e13, e14)) | (e10 = op1(e12, e14)) | (e10 = op1(e11, e14)) | (e10 = op1(e10, e14))) & ((e10 = op1(e14, e14)) | (e10 = op1(e14, e13)) | (e10 = op1(e14, e12)) | (e10 = op1(e14, e11)) | (e10 = op1(e14, e10))) & ((e14 = op1(e14, e13)) | (e14 = op1(e13, e13)) | (e14 = op1(e12, e13)) | (e14 = op1(e11, e13)) | (e14 = op1(e10, e13))) & ((e14 = op1(e13, e14)) | (e14 = op1(e13, e13)) | (e14 = op1(e13, e12)) | (e14 = op1(e13, e11)) | (e14 = op1(e13, e10))) & ((e13 = op1(e14, e13)) | (e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e14)) | (e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e14, e13)) | (e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e14)) | (e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e14, e13)) | (e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e14)) | (e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e14, e13)) | (e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e14)) | (e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e14 = op1(e14, e12)) | (e14 = op1(e13, e12)) | (e14 = op1(e12, e12)) | (e14 = op1(e11, e12)) | (e14 = op1(e10, e12))) & ((e14 = op1(e12, e14)) | (e14 = op1(e12, e13)) | (e14 = op1(e12, e12)) | (e14 = op1(e12, e11)) | (e14 = op1(e12, e10))) & ((e13 = op1(e14, e12)) | (e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e14)) | (e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e14, e12)) | (e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e14)) | (e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e14, e12)) | (e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e14)) | (e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e14, e12)) | (e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e14)) | (e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e14 = op1(e14, e11)) | (e14 = op1(e13, e11)) | (e14 = op1(e12, e11)) | (e14 = op1(e11, e11)) | (e14 = op1(e10, e11))) & ((e14 = op1(e11, e14)) | (e14 = op1(e11, e13)) | (e14 = op1(e11, e12)) | (e14 = op1(e11, e11)) | (e14 = op1(e11, e10))) & ((e13 = op1(e14, e11)) | (e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e14)) | (e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e14, e11)) | (e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e14)) | (e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e14, e11)) | (e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e14)) | (e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e14, e11)) | (e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e14)) | (e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e14 = op1(e14, e10)) | (e14 = op1(e13, e10)) | (e14 = op1(e12, e10)) | (e14 = op1(e11, e10)) | (op1(e10, e10) = e14)) & ((e14 = op1(e10, e14)) | (e14 = op1(e10, e13)) | (e14 = op1(e10, e12)) | (e14 = op1(e10, e11)) | (op1(e10, e10) = e14)) & ((e13 = op1(e14, e10)) | (e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e14)) | (e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e14, e10)) | (e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e14)) | (e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e14, e10)) | (e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e14)) | (e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e14, e10)) | (e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e14)) | (e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax2)).
fof(f1865, plain, (spl80_96 | spl80_91 | spl80_86 | spl80_81 | spl80_76), inference(avatar_split_clause, [], [f235, f1617, f1638, f1659, f1680, f1701])).
fof(f235, plain, ((e10 = op1(e11, e14)) | (e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1863, plain, (spl80_97 | spl80_92 | spl80_87 | spl80_82 | spl80_77), inference(avatar_split_clause, [], [f237, f1621, f1642, f1663, f1684, f1705])).
fof(f237, plain, ((e11 = op1(e11, e14)) | (e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1842, plain, (spl80_107 | spl80_82 | spl80_57 | spl80_32 | spl80_7), inference(avatar_split_clause, [], [f258, f1327, f1432, f1537, f1642, f1747])).
fof(f258, plain, ((e11 = op1(e14, e13)) | (e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1838, plain, (spl80_109 | spl80_84 | spl80_59 | spl80_34 | spl80_9), inference(avatar_split_clause, [], [f262, f1335, f1440, f1545, f1650, f1755])).
fof(f262, plain, ((e13 = op1(e14, e13)) | (e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1837, plain, (spl80_50 | spl80_45 | spl80_40 | spl80_35 | spl80_30), inference(avatar_split_clause, [], [f263, f1423, f1444, f1465, f1486, f1507])).
fof(f263, plain, ((e14 = op1(e13, e14)) | (e14 = op1(e13, e13)) | (e14 = op1(e13, e12)) | (e14 = op1(e13, e11)) | (e14 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1828, plain, (spl80_104 | spl80_79 | spl80_54 | spl80_29 | spl80_4), inference(avatar_split_clause, [], [f272, f1314, f1419, f1524, f1629, f1734])).
fof(f272, plain, ((e13 = op1(e14, e14)) | (e13 = op1(e13, e14)) | (e13 = op1(e12, e14)) | (e13 = op1(e11, e14)) | (e13 = op1(e10, e14))), inference(cnf_transformation, [], [f2])).
fof(f1826, plain, (spl80_105 | spl80_80 | spl80_55 | spl80_30 | spl80_5), inference(avatar_split_clause, [], [f274, f1318, f1423, f1528, f1633, f1738])).
fof(f274, plain, ((e14 = op1(e14, e14)) | (e14 = op1(e13, e14)) | (e14 = op1(e12, e14)) | (e14 = op1(e11, e14)) | (e14 = op1(e10, e14))), inference(cnf_transformation, [], [f2])).
fof(f1594, plain, (spl80_66 | spl80_67 | spl80_68 | spl80_69 | spl80_70), inference(avatar_split_clause, [], [f211, f1591, f1587, f1583, f1579, f1575])).
fof(f211, plain, ((e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e14 = op1(e14, e14)) | (e13 = op1(e14, e14)) | (e12 = op1(e14, e14)) | (e11 = op1(e14, e14)) | (e10 = op1(e14, e14))) & ((e14 = op1(e14, e13)) | (e13 = op1(e14, e13)) | (e12 = op1(e14, e13)) | (e11 = op1(e14, e13)) | (e10 = op1(e14, e13))) & ((e14 = op1(e14, e12)) | (e13 = op1(e14, e12)) | (e12 = op1(e14, e12)) | (e11 = op1(e14, e12)) | (e10 = op1(e14, e12))) & ((e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))) & ((e14 = op1(e14, e10)) | (e13 = op1(e14, e10)) | (e12 = op1(e14, e10)) | (e11 = op1(e14, e10)) | (e10 = op1(e14, e10))) & ((e14 = op1(e13, e14)) | (e13 = op1(e13, e14)) | (e12 = op1(e13, e14)) | (e11 = op1(e13, e14)) | (e10 = op1(e13, e14))) & ((e14 = op1(e13, e13)) | (e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e14 = op1(e13, e11)) | (e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e14 = op1(e13, e10)) | (e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e14 = op1(e12, e14)) | (e13 = op1(e12, e14)) | (e12 = op1(e12, e14)) | (e11 = op1(e12, e14)) | (e10 = op1(e12, e14))) & ((e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e14 = op1(e12, e12)) | (e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e14 = op1(e12, e10)) | (e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))) & ((e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e14 = op1(e11, e12)) | (e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e14 = op1(e11, e11)) | (e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e14 = op1(e11, e10)) | (e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))) & ((e14 = op1(e10, e13)) | (e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e14 = op1(e10, e12)) | (e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG169+1.p', ax1)).
fof(f1573, plain, (spl80_61 | spl80_62 | spl80_63 | spl80_64 | spl80_65), inference(avatar_split_clause, [], [f212, f1570, f1566, f1562, f1558, f1554])).
fof(f212, plain, ((e14 = op1(e12, e12)) | (e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))), inference(cnf_transformation, [], [f1])).