fof(f25677, plain, $false, inference(avatar_sat_refutation, [], [f1686, f1715, f1889, f1918, f1947, f2034, f2063, f2266, f2353, f2527, f2585, f2701, f2962, f3023, f3035, f3044, f3048, f3054, f3056, f3062, f3067, f3070, f3073, f3074, f3078, f3081, f3082, f3087, f3094, f3099, f3117, f3118, f3205, f3350, f3408, f3553, f3582, f3640, f4046, f4104, f4133, f4162, f4220, f4541, f4542, f4543, f4552, f4563, f4567, f4573, f4575, f4581, f4584, f4586, f4589, f4600, f4601, f4602, f4604, f4611, f4618, f4621, f4624, f4628, f4630, f4636, f4637, f4641, f4642, f4643, f4644, f4645, f4647, f4648, f4649, f4650, f4651, f4655, f4657, f4658, f4661, f4665, f4667, f4668, f4671, f5063, f5089, f5115, f5145, f5174, f5200, f6867, f8825, f8826, f8827, f8828, f8831, f8847, f8865, f8967, f9041, f9066, f9073, f9090, f9106, f9108, f9110, f9111, f9147, f9174, f9231, f9242, f9284, f9307, f9326, f9350, f9359, f9473, f9578, f9772, f9773, f9775, f9776, f9777, f9793, f9931, f10000, f10023, f10064, f10077, f10079, f10080, f10081, f10176, f10208, f10209, f10210, f10220, f10255, f10276, f10293, f10306, f10319, f10320, f10321, f10322, f10333, f10366, f10439, f10443, f10561, f10646, f10722, f10728, f10734, f10737, f10764, f10766, f10780, f10791, f10800, f10804, f10843, f10850, f10858, f10905, f10907, f10909, f10944, f11166, f11193, f11225, f11236, f11244, f11259, f11261, f11263, f11277, f11279, f11280, f11284, f11285, f11300, f11303, f11304, f11305, f11306, f11315, f11316, f11375, f11426, f11446, f11610, f11737, f11843, f11874, f11878, f12099, f12141, f12299, f12336, f12345, f12413, f12442, f12475, f12491, f12530, f12534, f12565, f12569, f12575, f12595, f12621, f12644, f12665, f12725, f12735, f12756, f12760, f12793, f12800, f12826, f12832, f12849, f12877, f12921, f12944, f12947, f13096, f13119, f13131, f13185, f13232, f13266, f13282, f13295, f13299, f13311, f13331, f13336, f13359, f13404, f13426, f13499, f13504, f13543, f13545, f13551, f13587, f13623, f13663, f13671, f13674, f13694, f13761, f13775, f13796, f13799, f13840, f13849, f13879, f13890, f13935, f14055, f14083, f14086, f14093, f14132, f14199, f14213, f14243, f14262, f14292, f14338, f14396, f14409, f14418, f14429, f14480, f14512, f14548, f14578, f14588, f14590, f14671, f14703, f14708, f14772, f14787, f14803, f14816, f14826, f14827, f14879, f14922, f14982, f15014, f15067, f15081, f15084, f15126, f15173, f15192, f15204, f15207, f15256, f15261, f15317, f15358, f15362, f15443, f15510, f15552, f15581, f15598, f15609, f15616, f15642, f15647, f15682, f15690, f15704, f15715, f15854, f15885, f16016, f16032, f16068, f16074, f16092, f16129, f16158, f16181, f16216, f16230, f16239, f16307, f16355, f16454, f16488, f16516, f16538, f16690, f16727, f16737, f16835, f16838, f16849, f16860, f16931, f16964, f16997, f17043, f17126, f17134, f17205, f17276, f17310, f17444, f17475, f17477, f17509, f17553, f17556, f17562, f17585, f17622, f17705, f17746, f17871, f17929, f17979, f18015, f18055, f18175, f18218, f18219, f18230, f18386, f18603, f18614, f18634, f18684, f18700, f18711, f18730, f18770, f18923, f18957, f18973, f18993, f19002, f19098, f19136, f19166, f19187, f19196, f19249, f19269, f19272, f19277, f19285, f19288, f19308, f19318, f19324, f19379, f19494, f19564, f19573, f19606, f19706, f19788, f19798, f19801, f19850, f19865, f19874, f19889, f19928, f19929, f19932, f19956, f20071, f20083, f20096, f20212, f20280, f20317, f20377, f20456, f20460, f20462, f20470, f20471, f20472, f20492, f20505, f20585, f20612, f20657, f20727, f20828, f20906, f20973, f20988, f20993, f21001, f21042, f21093, f21238, f21249, f21268, f21476, f21514, f21543, f21593, f21649, f21720, f21800, f21863, f21923, f21977, f22030, f22085, f22146, f22207, f22268, f22329, f22394, f22459, f22520, f22579, f22642, f22707, f22767, f22832, f22893, f22954, f23020, f23086, f23153, f23215, f23276, f23338, f23400, f23463, f23530, f23594, f23659, f23719, f23780, f23845, f23903, f23956, f24009, f24058, f24109, f24163, f24210, f24265, f24310, f24419, f24547, f24776, f24877, f24962, f25071, f25094, f25107, f25134, f25158, f25259, f25317, f25410, f25439])).
fof(f25439, plain, (spl42_31 | ~ spl42_84 | ~ spl42_230), inference(avatar_contradiction_clause, [], [f25438])).
fof(f25438, plain, ($false | (spl42_31 | ~ spl42_84 | ~ spl42_230)), inference(subsumption_resolution, [], [f25421, f1726])).
fof(f1726, plain, (~ (e12 = op1(e16, e12)) | spl42_31), inference(avatar_component_clause, [], [f1725])).
fof(f1725, plain, (spl42_31 <=> (e12 = op1(e16, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_31])])).
fof(f25421, plain, ((e12 = op1(e16, e12)) | (~ spl42_84 | ~ spl42_230)), inference(backward_demodulation, [], [f24565, f1946])).
fof(f1946, plain, ((e16 = op1(e15, e12)) | ~ spl42_84), inference(avatar_component_clause, [], [f1944])).
fof(f1944, plain, (spl42_84 <=> (e16 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_84])])).
fof(f24565, plain, ((e12 = op1(op1(e15, e12), e12)) | ~ spl42_230), inference(forward_demodulation, [], [f274, f2551])).
fof(f2551, plain, ((e15 = op1(e12, e12)) | ~ spl42_230), inference(avatar_component_clause, [], [f2549])).
fof(f2549, plain, (spl42_230 <=> (e15 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_230])])).
fof(f274, plain, (e12 = op1(op1(op1(e12, e12), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, ((e16 = op1(op1(op1(e16, e16), e16), e16)) & (e16 = op1(op1(op1(e15, e16), e15), e15)) & (e16 = op1(op1(op1(e14, e16), e14), e14)) & (e16 = op1(op1(op1(e13, e16), e13), e13)) & (e16 = op1(op1(op1(e12, e16), e12), e12)) & (e16 = op1(op1(op1(e11, e16), e11), e11)) & (e16 = op1(op1(op1(e10, e16), e10), e10)) & (e15 = op1(op1(op1(e16, e15), e16), e16)) & (e15 = op1(op1(op1(e15, e15), e15), e15)) & (e15 = op1(op1(op1(e14, e15), e14), e14)) & (e15 = op1(op1(op1(e13, e15), e13), e13)) & (e15 = op1(op1(op1(e12, e15), e12), e12)) & (e15 = op1(op1(op1(e11, e15), e11), e11)) & (e15 = op1(op1(op1(e10, e15), e10), e10)) & (e14 = op1(op1(op1(e16, e14), e16), e16)) & (e14 = op1(op1(op1(e15, e14), e15), e15)) & (e14 = op1(op1(op1(e14, e14), e14), e14)) & (e14 = op1(op1(op1(e13, e14), e13), e13)) & (e14 = op1(op1(op1(e12, e14), e12), e12)) & (e14 = op1(op1(op1(e11, e14), e11), e11)) & (e14 = op1(op1(op1(e10, e14), e10), e10)) & (e13 = op1(op1(op1(e16, e13), e16), e16)) & (e13 = op1(op1(op1(e15, e13), e15), e15)) & (e13 = op1(op1(op1(e14, e13), e14), e14)) & (e13 = op1(op1(op1(e13, e13), e13), e13)) & (e13 = op1(op1(op1(e12, e13), e12), e12)) & (e13 = op1(op1(op1(e11, e13), e11), e11)) & (e13 = op1(op1(op1(e10, e13), e10), e10)) & (e12 = op1(op1(op1(e16, e12), e16), e16)) & (e12 = op1(op1(op1(e15, e12), e15), e15)) & (e12 = op1(op1(op1(e14, e12), e14), e14)) & (e12 = op1(op1(op1(e13, e12), e13), e13)) & (e12 = op1(op1(op1(e12, e12), e12), e12)) & (e12 = op1(op1(op1(e11, e12), e11), e11)) & (e12 = op1(op1(op1(e10, e12), e10), e10)) & (e11 = op1(op1(op1(e16, e11), e16), e16)) & (e11 = op1(op1(op1(e15, e11), e15), e15)) & (e11 = op1(op1(op1(e14, e11), e14), e14)) & (e11 = op1(op1(op1(e13, e11), e13), e13)) & (e11 = op1(op1(op1(e12, e11), e12), e12)) & (e11 = op1(op1(op1(e11, e11), e11), e11)) & (e11 = op1(op1(op1(e10, e11), e10), e10)) & (e10 = op1(op1(op1(e16, e10), e16), e16)) & (e10 = op1(op1(op1(e15, e10), e15), e15)) & (e10 = op1(op1(op1(e14, e10), e14), e14)) & (e10 = op1(op1(op1(e13, e10), e13), e13)) & (e10 = op1(op1(op1(e12, e10), e12), e12)) & (e10 = op1(op1(op1(e11, e10), e11), e11)) & (e10 = op1(op1(op1(e10, e10), e10), e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax3)).
fof(f25410, plain, (~ spl42_87 | ~ spl42_89), inference(avatar_contradiction_clause, [], [f25409])).
fof(f25409, plain, ($false | (~ spl42_87 | ~ spl42_89)), inference(subsumption_resolution, [], [f25408, f1103])).
fof(f1103, plain, ~ (e12 = e14), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (~ (e15 = e16) & ~ (e14 = e16) & ~ (e14 = e15) & ~ (e13 = e16) & ~ (e13 = e15) & ~ (e13 = e14) & ~ (e12 = e16) & ~ (e12 = e15) & ~ (e12 = e14) & ~ (e12 = e13) & ~ (e11 = e16) & ~ (e11 = e15) & ~ (e11 = e14) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e16) & ~ (e10 = e15) & ~ (e10 = e14) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax9)).
fof(f25408, plain, ((e12 = e14) | (~ spl42_87 | ~ spl42_89)), inference(backward_demodulation, [], [f1967, f1959])).
fof(f1959, plain, ((e12 = op1(e15, e11)) | ~ spl42_87), inference(avatar_component_clause, [], [f1957])).
fof(f1957, plain, (spl42_87 <=> (e12 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_87])])).
fof(f1967, plain, ((e14 = op1(e15, e11)) | ~ spl42_89), inference(avatar_component_clause, [], [f1965])).
fof(f1965, plain, (spl42_89 <=> (e14 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_89])])).
fof(f25317, plain, (spl42_9 | ~ spl42_89 | ~ spl42_112), inference(avatar_contradiction_clause, [], [f25316])).
fof(f25316, plain, ($false | (spl42_9 | ~ spl42_89 | ~ spl42_112)), inference(subsumption_resolution, [], [f25304, f1635])).
fof(f1635, plain, (~ (e11 = op1(e16, e15)) | spl42_9), inference(avatar_component_clause, [], [f1634])).
fof(f1634, plain, (spl42_9 <=> (e11 = op1(e16, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_9])])).
fof(f25304, plain, ((e11 = op1(e16, e15)) | (~ spl42_89 | ~ spl42_112)), inference(backward_demodulation, [], [f13300, f2062])).
fof(f2062, plain, ((e16 = op1(e14, e15)) | ~ spl42_112), inference(avatar_component_clause, [], [f2060])).
fof(f2060, plain, (spl42_112 <=> (e16 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_112])])).
fof(f13300, plain, ((e11 = op1(op1(e14, e15), e15)) | ~ spl42_89), inference(forward_demodulation, [], [f270, f1967])).
fof(f270, plain, (e11 = op1(op1(op1(e15, e11), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f25259, plain, (~ spl42_144 | ~ spl42_145), inference(avatar_contradiction_clause, [], [f25258])).
fof(f25258, plain, ($false | (~ spl42_144 | ~ spl42_145)), inference(subsumption_resolution, [], [f25257, f1106])).
fof(f1106, plain, ~ (e13 = e14), inference(cnf_transformation, [], [f9])).
fof(f25257, plain, ((e13 = e14) | (~ spl42_144 | ~ spl42_145)), inference(backward_demodulation, [], [f2199, f2195])).
fof(f2195, plain, ((e13 = op1(e14, e10)) | ~ spl42_144), inference(avatar_component_clause, [], [f2193])).
fof(f2193, plain, (spl42_144 <=> (e13 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_144])])).
fof(f2199, plain, ((e14 = op1(e14, e10)) | ~ spl42_145), inference(avatar_component_clause, [], [f2197])).
fof(f2197, plain, (spl42_145 <=> (e14 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_145])])).
fof(f25158, plain, (~ spl42_155 | ~ spl42_190), inference(avatar_split_clause, [], [f25138, f2384, f2239])).
fof(f2239, plain, (spl42_155 <=> (e10 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_155])])).
fof(f2384, plain, (spl42_190 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_190])])).
fof(f25138, plain, (~ (e10 = op1(e13, e15)) | ~ spl42_190), inference(backward_demodulation, [], [f717, f2386])).
fof(f2386, plain, ((e10 = op1(e13, e10)) | ~ spl42_190), inference(avatar_component_clause, [], [f2384])).
fof(f717, plain, ~ (op1(e13, e10) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (op1(e16, e15) = op1(e16, e16)) & ~ (op1(e16, e14) = op1(e16, e16)) & ~ (op1(e16, e14) = op1(e16, e15)) & ~ (op1(e16, e13) = op1(e16, e16)) & ~ (op1(e16, e13) = op1(e16, e15)) & ~ (op1(e16, e13) = op1(e16, e14)) & ~ (op1(e16, e12) = op1(e16, e16)) & ~ (op1(e16, e12) = op1(e16, e15)) & ~ (op1(e16, e12) = op1(e16, e14)) & ~ (op1(e16, e12) = op1(e16, e13)) & ~ (op1(e16, e11) = op1(e16, e16)) & ~ (op1(e16, e11) = op1(e16, e15)) & ~ (op1(e16, e11) = op1(e16, e14)) & ~ (op1(e16, e11) = op1(e16, e13)) & ~ (op1(e16, e11) = op1(e16, e12)) & ~ (op1(e16, e10) = op1(e16, e16)) & ~ (op1(e16, e10) = op1(e16, e15)) & ~ (op1(e16, e10) = op1(e16, e14)) & ~ (op1(e16, e10) = op1(e16, e13)) & ~ (op1(e16, e10) = op1(e16, e12)) & ~ (op1(e16, e10) = op1(e16, e11)) & ~ (op1(e15, e15) = op1(e15, e16)) & ~ (op1(e15, e14) = op1(e15, e16)) & ~ (op1(e15, e14) = op1(e15, e15)) & ~ (op1(e15, e13) = op1(e15, e16)) & ~ (op1(e15, e13) = op1(e15, e15)) & ~ (op1(e15, e13) = op1(e15, e14)) & ~ (op1(e15, e12) = op1(e15, e16)) & ~ (op1(e15, e12) = op1(e15, e15)) & ~ (op1(e15, e12) = op1(e15, e14)) & ~ (op1(e15, e12) = op1(e15, e13)) & ~ (op1(e15, e11) = op1(e15, e16)) & ~ (op1(e15, e11) = op1(e15, e15)) & ~ (op1(e15, e11) = op1(e15, e14)) & ~ (op1(e15, e11) = op1(e15, e13)) & ~ (op1(e15, e11) = op1(e15, e12)) & ~ (op1(e15, e10) = op1(e15, e16)) & ~ (op1(e15, e10) = op1(e15, e15)) & ~ (op1(e15, e10) = op1(e15, e14)) & ~ (op1(e15, e10) = op1(e15, e13)) & ~ (op1(e15, e10) = op1(e15, e12)) & ~ (op1(e15, e10) = op1(e15, e11)) & ~ (op1(e14, e15) = op1(e14, e16)) & ~ (op1(e14, e14) = op1(e14, e16)) & ~ (op1(e14, e14) = op1(e14, e15)) & ~ (op1(e14, e13) = op1(e14, e16)) & ~ (op1(e14, e13) = op1(e14, e15)) & ~ (op1(e14, e13) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e16)) & ~ (op1(e14, e12) = op1(e14, e15)) & ~ (op1(e14, e12) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e16)) & ~ (op1(e14, e11) = op1(e14, e15)) & ~ (op1(e14, e11) = op1(e14, e14)) & ~ (op1(e14, e11) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e16)) & ~ (op1(e14, e10) = op1(e14, e15)) & ~ (op1(e14, e10) = op1(e14, e14)) & ~ (op1(e14, e10) = op1(e14, e13)) & ~ (op1(e14, e10) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e11)) & ~ (op1(e13, e15) = op1(e13, e16)) & ~ (op1(e13, e14) = op1(e13, e16)) & ~ (op1(e13, e14) = op1(e13, e15)) & ~ (op1(e13, e13) = op1(e13, e16)) & ~ (op1(e13, e13) = op1(e13, e15)) & ~ (op1(e13, e13) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e16)) & ~ (op1(e13, e12) = op1(e13, e15)) & ~ (op1(e13, e12) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e16)) & ~ (op1(e13, e11) = op1(e13, e15)) & ~ (op1(e13, e11) = op1(e13, e14)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e16)) & ~ (op1(e13, e10) = op1(e13, e15)) & ~ (op1(e13, e10) = op1(e13, e14)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e15) = op1(e12, e16)) & ~ (op1(e12, e14) = op1(e12, e16)) & ~ (op1(e12, e14) = op1(e12, e15)) & ~ (op1(e12, e13) = op1(e12, e16)) & ~ (op1(e12, e13) = op1(e12, e15)) & ~ (op1(e12, e13) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e16)) & ~ (op1(e12, e12) = op1(e12, e15)) & ~ (op1(e12, e12) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e16)) & ~ (op1(e12, e11) = op1(e12, e15)) & ~ (op1(e12, e11) = op1(e12, e14)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e16)) & ~ (op1(e12, e10) = op1(e12, e15)) & ~ (op1(e12, e10) = op1(e12, e14)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e15) = op1(e11, e16)) & ~ (op1(e11, e14) = op1(e11, e16)) & ~ (op1(e11, e14) = op1(e11, e15)) & ~ (op1(e11, e13) = op1(e11, e16)) & ~ (op1(e11, e13) = op1(e11, e15)) & ~ (op1(e11, e13) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e16)) & ~ (op1(e11, e12) = op1(e11, e15)) & ~ (op1(e11, e12) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e16)) & ~ (op1(e11, e11) = op1(e11, e15)) & ~ (op1(e11, e11) = op1(e11, e14)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e16)) & ~ (op1(e11, e10) = op1(e11, e15)) & ~ (op1(e11, e10) = op1(e11, e14)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e15) = op1(e10, e16)) & ~ (op1(e10, e14) = op1(e10, e16)) & ~ (op1(e10, e14) = op1(e10, e15)) & ~ (op1(e10, e13) = op1(e10, e16)) & ~ (op1(e10, e13) = op1(e10, e15)) & ~ (op1(e10, e13) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e16)) & ~ (op1(e10, e12) = op1(e10, e15)) & ~ (op1(e10, e12) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e16)) & ~ (op1(e10, e11) = op1(e10, e15)) & ~ (op1(e10, e11) = op1(e10, e14)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e16)) & ~ (op1(e10, e10) = op1(e10, e15)) & ~ (op1(e10, e10) = op1(e10, e14)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e15, e16) = op1(e16, e16)) & ~ (op1(e14, e16) = op1(e16, e16)) & ~ (op1(e14, e16) = op1(e15, e16)) & ~ (op1(e13, e16) = op1(e16, e16)) & ~ (op1(e13, e16) = op1(e15, e16)) & ~ (op1(e13, e16) = op1(e14, e16)) & ~ (op1(e12, e16) = op1(e16, e16)) & ~ (op1(e12, e16) = op1(e15, e16)) & ~ (op1(e12, e16) = op1(e14, e16)) & ~ (op1(e12, e16) = op1(e13, e16)) & ~ (op1(e11, e16) = op1(e16, e16)) & ~ (op1(e11, e16) = op1(e15, e16)) & ~ (op1(e11, e16) = op1(e14, e16)) & ~ (op1(e11, e16) = op1(e13, e16)) & ~ (op1(e11, e16) = op1(e12, e16)) & ~ (op1(e10, e16) = op1(e16, e16)) & ~ (op1(e10, e16) = op1(e15, e16)) & ~ (op1(e10, e16) = op1(e14, e16)) & ~ (op1(e10, e16) = op1(e13, e16)) & ~ (op1(e10, e16) = op1(e12, e16)) & ~ (op1(e10, e16) = op1(e11, e16)) & ~ (op1(e15, e15) = op1(e16, e15)) & ~ (op1(e14, e15) = op1(e16, e15)) & ~ (op1(e14, e15) = op1(e15, e15)) & ~ (op1(e13, e15) = op1(e16, e15)) & ~ (op1(e13, e15) = op1(e15, e15)) & ~ (op1(e13, e15) = op1(e14, e15)) & ~ (op1(e12, e15) = op1(e16, e15)) & ~ (op1(e12, e15) = op1(e15, e15)) & ~ (op1(e12, e15) = op1(e14, e15)) & ~ (op1(e12, e15) = op1(e13, e15)) & ~ (op1(e11, e15) = op1(e16, e15)) & ~ (op1(e11, e15) = op1(e15, e15)) & ~ (op1(e11, e15) = op1(e14, e15)) & ~ (op1(e11, e15) = op1(e13, e15)) & ~ (op1(e11, e15) = op1(e12, e15)) & ~ (op1(e10, e15) = op1(e16, e15)) & ~ (op1(e10, e15) = op1(e15, e15)) & ~ (op1(e10, e15) = op1(e14, e15)) & ~ (op1(e10, e15) = op1(e13, e15)) & ~ (op1(e10, e15) = op1(e12, e15)) & ~ (op1(e10, e15) = op1(e11, e15)) & ~ (op1(e15, e14) = op1(e16, e14)) & ~ (op1(e14, e14) = op1(e16, e14)) & ~ (op1(e14, e14) = op1(e15, e14)) & ~ (op1(e13, e14) = op1(e16, e14)) & ~ (op1(e13, e14) = op1(e15, e14)) & ~ (op1(e13, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e16, e14)) & ~ (op1(e12, e14) = op1(e15, e14)) & ~ (op1(e12, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e16, e14)) & ~ (op1(e11, e14) = op1(e15, e14)) & ~ (op1(e11, e14) = op1(e14, e14)) & ~ (op1(e11, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e16, e14)) & ~ (op1(e10, e14) = op1(e15, e14)) & ~ (op1(e10, e14) = op1(e14, e14)) & ~ (op1(e10, e14) = op1(e13, e14)) & ~ (op1(e10, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e11, e14)) & ~ (op1(e15, e13) = op1(e16, e13)) & ~ (op1(e14, e13) = op1(e16, e13)) & ~ (op1(e14, e13) = op1(e15, e13)) & ~ (op1(e13, e13) = op1(e16, e13)) & ~ (op1(e13, e13) = op1(e15, e13)) & ~ (op1(e13, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e16, e13)) & ~ (op1(e12, e13) = op1(e15, e13)) & ~ (op1(e12, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e16, e13)) & ~ (op1(e11, e13) = op1(e15, e13)) & ~ (op1(e11, e13) = op1(e14, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e16, e13)) & ~ (op1(e10, e13) = op1(e15, e13)) & ~ (op1(e10, e13) = op1(e14, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e15, e12) = op1(e16, e12)) & ~ (op1(e14, e12) = op1(e16, e12)) & ~ (op1(e14, e12) = op1(e15, e12)) & ~ (op1(e13, e12) = op1(e16, e12)) & ~ (op1(e13, e12) = op1(e15, e12)) & ~ (op1(e13, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e16, e12)) & ~ (op1(e12, e12) = op1(e15, e12)) & ~ (op1(e12, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e16, e12)) & ~ (op1(e11, e12) = op1(e15, e12)) & ~ (op1(e11, e12) = op1(e14, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e16, e12)) & ~ (op1(e10, e12) = op1(e15, e12)) & ~ (op1(e10, e12) = op1(e14, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e15, e11) = op1(e16, e11)) & ~ (op1(e14, e11) = op1(e16, e11)) & ~ (op1(e14, e11) = op1(e15, e11)) & ~ (op1(e13, e11) = op1(e16, e11)) & ~ (op1(e13, e11) = op1(e15, e11)) & ~ (op1(e13, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e16, e11)) & ~ (op1(e12, e11) = op1(e15, e11)) & ~ (op1(e12, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e16, e11)) & ~ (op1(e11, e11) = op1(e15, e11)) & ~ (op1(e11, e11) = op1(e14, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e16, e11)) & ~ (op1(e10, e11) = op1(e15, e11)) & ~ (op1(e10, e11) = op1(e14, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e15, e10) = op1(e16, e10)) & ~ (op1(e14, e10) = op1(e16, e10)) & ~ (op1(e14, e10) = op1(e15, e10)) & ~ (op1(e13, e10) = op1(e16, e10)) & ~ (op1(e13, e10) = op1(e15, e10)) & ~ (op1(e13, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e16, e10)) & ~ (op1(e12, e10) = op1(e15, e10)) & ~ (op1(e12, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e16, e10)) & ~ (op1(e11, e10) = op1(e15, e10)) & ~ (op1(e11, e10) = op1(e14, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e16, e10)) & ~ (op1(e10, e10) = op1(e15, e10)) & ~ (op1(e10, e10) = op1(e14, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax7)).
fof(f25134, plain, (~ spl42_201 | ~ spl42_203), inference(avatar_contradiction_clause, [], [f25133])).
fof(f25133, plain, ($false | (~ spl42_201 | ~ spl42_203)), inference(subsumption_resolution, [], [f25132, f1110])).
fof(f1110, plain, ~ (e14 = e16), inference(cnf_transformation, [], [f9])).
fof(f25132, plain, ((e14 = e16) | (~ spl42_201 | ~ spl42_203)), inference(forward_demodulation, [], [f2439, f2431])).
fof(f2431, plain, ((e14 = op1(e12, e16)) | ~ spl42_201), inference(avatar_component_clause, [], [f2429])).
fof(f2429, plain, (spl42_201 <=> (e14 = op1(e12, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_201])])).
fof(f2439, plain, ((e16 = op1(e12, e16)) | ~ spl42_203), inference(avatar_component_clause, [], [f2437])).
fof(f2437, plain, (spl42_203 <=> (e16 = op1(e12, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_203])])).
fof(f25107, plain, (~ spl42_223 | ~ spl42_230), inference(avatar_contradiction_clause, [], [f25106])).
fof(f25106, plain, ($false | (~ spl42_223 | ~ spl42_230)), inference(subsumption_resolution, [], [f25105, f1102])).
fof(f1102, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f9])).
fof(f25105, plain, ((e12 = e13) | (~ spl42_223 | ~ spl42_230)), inference(forward_demodulation, [], [f25095, f24565])).
fof(f25095, plain, ((e13 = op1(op1(e15, e12), e12)) | ~ spl42_223), inference(backward_demodulation, [], [f281, f2522])).
fof(f2522, plain, ((e15 = op1(e12, e13)) | ~ spl42_223), inference(avatar_component_clause, [], [f2520])).
fof(f2520, plain, (spl42_223 <=> (e15 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_223])])).
fof(f281, plain, (e13 = op1(op1(op1(e12, e13), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f25094, plain, (~ spl42_226 | ~ spl42_230), inference(avatar_contradiction_clause, [], [f25093])).
fof(f25093, plain, ($false | (~ spl42_226 | ~ spl42_230)), inference(subsumption_resolution, [], [f25092, f1100])).
fof(f1100, plain, ~ (e11 = e15), inference(cnf_transformation, [], [f9])).
fof(f25092, plain, ((e11 = e15) | (~ spl42_226 | ~ spl42_230)), inference(backward_demodulation, [], [f2551, f2535])).
fof(f2535, plain, ((e11 = op1(e12, e12)) | ~ spl42_226), inference(avatar_component_clause, [], [f2533])).
fof(f2533, plain, (spl42_226 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_226])])).
fof(f25071, plain, (spl42_79 | ~ spl42_230 | ~ spl42_234), inference(avatar_contradiction_clause, [], [f25070])).
fof(f25070, plain, ($false | (spl42_79 | ~ spl42_230 | ~ spl42_234)), inference(subsumption_resolution, [], [f25069, f1925])).
fof(f1925, plain, (~ (e11 = op1(e15, e12)) | spl42_79), inference(avatar_component_clause, [], [f1924])).
fof(f1924, plain, (spl42_79 <=> (e11 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_79])])).
fof(f25069, plain, ((e11 = op1(e15, e12)) | (~ spl42_230 | ~ spl42_234)), inference(forward_demodulation, [], [f25054, f2551])).
fof(f25054, plain, ((e11 = op1(op1(e12, e12), e12)) | ~ spl42_234), inference(backward_demodulation, [], [f267, f2568])).
fof(f2568, plain, ((e12 = op1(e12, e11)) | ~ spl42_234), inference(avatar_component_clause, [], [f2566])).
fof(f2566, plain, (spl42_234 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_234])])).
fof(f267, plain, (e11 = op1(op1(op1(e12, e11), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f24962, plain, (~ spl42_58 | ~ spl42_258), inference(avatar_contradiction_clause, [], [f24961])).
fof(f24961, plain, ($false | (~ spl42_58 | ~ spl42_258)), inference(subsumption_resolution, [], [f24960, f1100])).
fof(f24960, plain, ((e11 = e15) | (~ spl42_58 | ~ spl42_258)), inference(forward_demodulation, [], [f24948, f1839])).
fof(f1839, plain, ((e11 = op1(e15, e15)) | ~ spl42_58), inference(avatar_component_clause, [], [f1837])).
fof(f1837, plain, (spl42_58 <=> (e11 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_58])])).
fof(f24948, plain, ((e15 = op1(e15, e15)) | (~ spl42_58 | ~ spl42_258)), inference(backward_demodulation, [], [f10430, f2667])).
fof(f2667, plain, ((e15 = op1(e11, e15)) | ~ spl42_258), inference(avatar_component_clause, [], [f2665])).
fof(f2665, plain, (spl42_258 <=> (e15 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_258])])).
fof(f10430, plain, ((e15 = op1(op1(e11, e15), e15)) | ~ spl42_58), inference(backward_demodulation, [], [f298, f1839])).
fof(f298, plain, (e15 = op1(op1(op1(e15, e15), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f24877, plain, (~ spl42_257 | ~ spl42_278), inference(avatar_split_clause, [], [f24866, f2748, f2661])).
fof(f2661, plain, (spl42_257 <=> (e14 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_257])])).
fof(f2748, plain, (spl42_278 <=> (e14 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_278])])).
fof(f24866, plain, (~ (e14 = op1(e11, e15)) | ~ spl42_278), inference(backward_demodulation, [], [f684, f2750])).
fof(f2750, plain, ((e14 = op1(e11, e12)) | ~ spl42_278), inference(avatar_component_clause, [], [f2748])).
fof(f684, plain, ~ (op1(e11, e12) = op1(e11, e15)), inference(cnf_transformation, [], [f7])).
fof(f24776, plain, (~ spl42_1 | ~ spl42_295), inference(avatar_split_clause, [], [f24758, f2819, f1601])).
fof(f1601, plain, (spl42_1 <=> (e10 = op1(e16, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_1])])).
fof(f2819, plain, (spl42_295 <=> (e10 = op1(e10, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_295])])).
fof(f24758, plain, (~ (e10 = op1(e16, e16)) | ~ spl42_295), inference(backward_demodulation, [], [f634, f2821])).
fof(f2821, plain, ((e10 = op1(e10, e16)) | ~ spl42_295), inference(avatar_component_clause, [], [f2819])).
fof(f634, plain, ~ (op1(e10, e16) = op1(e16, e16)), inference(cnf_transformation, [], [f7])).
fof(f24547, plain, (~ spl42_157 | ~ spl42_178), inference(avatar_split_clause, [], [f24546, f2334, f2247])).
fof(f2247, plain, (spl42_157 <=> (e12 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_157])])).
fof(f2334, plain, (spl42_178 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_178])])).
fof(f24546, plain, (~ (e12 = op1(e13, e15)) | ~ spl42_178), inference(forward_demodulation, [], [f726, f2336])).
fof(f2336, plain, ((e12 = op1(e13, e12)) | ~ spl42_178), inference(avatar_component_clause, [], [f2334])).
fof(f726, plain, ~ (op1(e13, e12) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f24419, plain, (~ spl42_43 | ~ spl42_48), inference(avatar_contradiction_clause, [], [f24418])).
fof(f24418, plain, ($false | (~ spl42_43 | ~ spl42_48)), inference(subsumption_resolution, [], [f24417, f1095])).
fof(f1095, plain, ~ (e10 = e15), inference(cnf_transformation, [], [f9])).
fof(f24417, plain, ((e10 = e15) | (~ spl42_43 | ~ spl42_48)), inference(backward_demodulation, [], [f1797, f1777])).
fof(f1777, plain, ((e10 = op1(e16, e10)) | ~ spl42_43), inference(avatar_component_clause, [], [f1775])).
fof(f1775, plain, (spl42_43 <=> (e10 = op1(e16, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_43])])).
fof(f1797, plain, ((e15 = op1(e16, e10)) | ~ spl42_48), inference(avatar_component_clause, [], [f1795])).
fof(f1795, plain, (spl42_48 <=> (e15 = op1(e16, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_48])])).
fof(f24310, plain, (~ spl42_211 | spl42_1045), inference(avatar_contradiction_clause, [], [f24309])).
fof(f24309, plain, ($false | (~ spl42_211 | spl42_1045)), inference(subsumption_resolution, [], [f24308, f4820])).
fof(f4820, plain, (e20 = h6(e10)), inference(forward_demodulation, [], [f4819, f4779])).
fof(f4779, plain, (e20 = op2(h5(e11), e24)), inference(backward_demodulation, [], [f4669, f1238])).
fof(f1238, plain, (op2(e24, e24) = h5(e11)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ((h5(e16) = op2(op2(op2(op2(e24, op2(e24, e24)), op2(e24, op2(e24, e24))), e24), op2(e24, op2(e24, e24)))) & (op2(e24, op2(e24, e24)) = h5(e14)) & (h5(e13) = op2(op2(op2(e24, op2(e24, e24)), op2(e24, op2(e24, e24))), e24)) & (op2(op2(e24, op2(e24, e24)), op2(e24, op2(e24, e24))) = h5(e12)) & (op2(e24, e24) = h5(e11)) & (h5(e10) = op2(op2(op2(e24, op2(e24, e24)), op2(e24, op2(e24, e24))), op2(e24, op2(e24, e24)))) & (e24 = h5(e15))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax20)).
fof(f4669, plain, (e20 = op2(op2(e24, e24), e24)), inference(forward_demodulation, [], [f1202, f1206])).
fof(f1206, plain, (e24 = op2(e25, op2(e25, e25))), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((e26 = op2(op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), e25), op2(e25, op2(e25, e25)))) & (e24 = op2(e25, op2(e25, e25))) & (e23 = op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), e25)) & (e22 = op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25)))) & (e21 = op2(e25, e25)) & (e20 = op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), op2(e25, op2(e25, e25))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax15)).
fof(f1202, plain, (e20 = op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), op2(e25, op2(e25, e25)))), inference(cnf_transformation, [], [f15])).
fof(f4819, plain, (h6(e10) = op2(h5(e11), e24)), inference(forward_demodulation, [], [f4818, f1238])).
fof(f4818, plain, (op2(op2(e24, e24), e24) = h6(e10)), inference(forward_demodulation, [], [f4817, f4816])).
fof(f4816, plain, (e24 = op2(e25, h6(e11))), inference(backward_demodulation, [], [f1206, f1245])).
fof(f1245, plain, (op2(e25, e25) = h6(e11)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ((op2(op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), e25), op2(e25, op2(e25, e25))) = h6(e16)) & (op2(e25, op2(e25, e25)) = h6(e14)) & (op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), e25) = h6(e13)) & (op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))) = h6(e12)) & (op2(e25, e25) = h6(e11)) & (op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), op2(e25, op2(e25, e25))) = h6(e10)) & (e25 = h6(e15))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax21)).
fof(f4817, plain, (h6(e10) = op2(op2(op2(e25, h6(e11)), op2(e25, h6(e11))), op2(e25, h6(e11)))), inference(forward_demodulation, [], [f1244, f1245])).
fof(f1244, plain, (op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), op2(e25, op2(e25, e25))) = h6(e10)), inference(cnf_transformation, [], [f21])).
fof(f24308, plain, (~ (e20 = h6(e10)) | (~ spl42_211 | spl42_1045)), inference(forward_demodulation, [], [f6802, f2473])).
fof(f2473, plain, ((e10 = op1(e12, e14)) | ~ spl42_211), inference(avatar_component_clause, [], [f2471])).
fof(f2471, plain, (spl42_211 <=> (e10 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_211])])).
fof(f6802, plain, (~ (e20 = h6(op1(e12, e14))) | spl42_1045), inference(avatar_component_clause, [], [f6800])).
fof(f6800, plain, (spl42_1045 <=> (e20 = h6(op1(e12, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1045])])).
fof(f24265, plain, (~ spl42_288 | ~ spl42_631 | ~ spl42_748 | spl42_1054), inference(avatar_contradiction_clause, [], [f24264])).
fof(f24264, plain, ($false | (~ spl42_288 | ~ spl42_631 | ~ spl42_748 | spl42_1054)), inference(subsumption_resolution, [], [f24263, f4311])).
fof(f4311, plain, ((e20 = op2(e21, e20)) | ~ spl42_631), inference(avatar_component_clause, [], [f4309])).
fof(f4309, plain, (spl42_631 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_631])])).
fof(f24263, plain, (~ (e20 = op2(e21, e20)) | (~ spl42_288 | ~ spl42_748 | spl42_1054)), inference(forward_demodulation, [], [f24262, f4820])).
fof(f24262, plain, (~ (op2(e21, e20) = h6(e10)) | (~ spl42_288 | ~ spl42_748 | spl42_1054)), inference(forward_demodulation, [], [f24261, f2792])).
fof(f2792, plain, ((e10 = op1(e11, e10)) | ~ spl42_288), inference(avatar_component_clause, [], [f2790])).
fof(f2790, plain, (spl42_288 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_288])])).
fof(f24261, plain, (~ (op2(e21, e20) = h6(op1(e11, e10))) | (~ spl42_748 | spl42_1054)), inference(forward_demodulation, [], [f6838, f5172])).
fof(f5172, plain, ((e21 = h6(e11)) | ~ spl42_748), inference(avatar_component_clause, [], [f5171])).
fof(f5171, plain, (spl42_748 <=> (e21 = h6(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_748])])).
fof(f6838, plain, (~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | spl42_1054), inference(avatar_component_clause, [], [f6836])).
fof(f6836, plain, (spl42_1054 <=> (h6(op1(e11, e10)) = op2(h6(e11), e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_1054])])).
fof(f24210, plain, (~ spl42_317 | ~ spl42_660 | ~ spl42_748 | spl42_1058), inference(avatar_contradiction_clause, [], [f24209])).
fof(f24209, plain, ($false | (~ spl42_317 | ~ spl42_660 | ~ spl42_748 | spl42_1058)), inference(subsumption_resolution, [], [f24208, f4431])).
fof(f4431, plain, ((e21 = op2(e20, e23)) | ~ spl42_660), inference(avatar_component_clause, [], [f4429])).
fof(f4429, plain, (spl42_660 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_660])])).
fof(f24208, plain, (~ (e21 = op2(e20, e23)) | (~ spl42_317 | ~ spl42_748 | spl42_1058)), inference(forward_demodulation, [], [f24207, f5172])).
fof(f24207, plain, (~ (op2(e20, e23) = h6(e11)) | (~ spl42_317 | spl42_1058)), inference(forward_demodulation, [], [f6854, f2912])).
fof(f2912, plain, ((e11 = op1(e10, e13)) | ~ spl42_317), inference(avatar_component_clause, [], [f2910])).
fof(f2910, plain, (spl42_317 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_317])])).
fof(f6854, plain, (~ (op2(e20, e23) = h6(op1(e10, e13))) | spl42_1058), inference(avatar_component_clause, [], [f6852])).
fof(f6852, plain, (spl42_1058 <=> (op2(e20, e23) = h6(op1(e10, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1058])])).
fof(f24163, plain, (~ spl42_245 | ~ spl42_588 | ~ spl42_743 | ~ spl42_1016 | spl42_1049), inference(avatar_contradiction_clause, [], [f24162])).
fof(f24162, plain, ($false | (~ spl42_245 | ~ spl42_588 | ~ spl42_743 | ~ spl42_1016 | spl42_1049)), inference(subsumption_resolution, [], [f24161, f4132])).
fof(f4132, plain, ((e26 = op2(e22, e20)) | ~ spl42_588), inference(avatar_component_clause, [], [f4130])).
fof(f4130, plain, (spl42_588 <=> (e26 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_588])])).
fof(f24161, plain, (~ (e26 = op2(e22, e20)) | (~ spl42_245 | ~ spl42_743 | ~ spl42_1016 | spl42_1049)), inference(forward_demodulation, [], [f24160, f6685])).
fof(f6685, plain, ((e26 = h6(e16)) | ~ spl42_1016), inference(avatar_component_clause, [], [f6684])).
fof(f6684, plain, (spl42_1016 <=> (e26 = h6(e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_1016])])).
fof(f24160, plain, (~ (op2(e22, e20) = h6(e16)) | (~ spl42_245 | ~ spl42_743 | spl42_1049)), inference(forward_demodulation, [], [f24159, f2613])).
fof(f2613, plain, ((e16 = op1(e12, e10)) | ~ spl42_245), inference(avatar_component_clause, [], [f2611])).
fof(f2611, plain, (spl42_245 <=> (e16 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_245])])).
fof(f24159, plain, (~ (op2(e22, e20) = h6(op1(e12, e10))) | (~ spl42_743 | spl42_1049)), inference(forward_demodulation, [], [f6818, f5143])).
fof(f5143, plain, ((e22 = h5(e11)) | ~ spl42_743), inference(avatar_component_clause, [], [f5142])).
fof(f5142, plain, (spl42_743 <=> (e22 = h5(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_743])])).
fof(f6818, plain, (~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | spl42_1049), inference(avatar_component_clause, [], [f6816])).
fof(f6816, plain, (spl42_1049 <=> (h6(op1(e12, e10)) = op2(h5(e11), e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_1049])])).
fof(f24109, plain, (~ spl42_178 | ~ spl42_521 | ~ spl42_743 | spl42_1040), inference(avatar_contradiction_clause, [], [f24108])).
fof(f24108, plain, ($false | (~ spl42_178 | ~ spl42_521 | ~ spl42_743 | spl42_1040)), inference(subsumption_resolution, [], [f24107, f3855])).
fof(f3855, plain, ((e22 = op2(e23, e22)) | ~ spl42_521), inference(avatar_component_clause, [], [f3853])).
fof(f3853, plain, (spl42_521 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_521])])).
fof(f24107, plain, (~ (e22 = op2(e23, e22)) | (~ spl42_178 | ~ spl42_743 | spl42_1040)), inference(forward_demodulation, [], [f24106, f11296])).
fof(f11296, plain, ((e22 = h6(e12)) | ~ spl42_743), inference(backward_demodulation, [], [f4802, f5143])).
fof(f4802, plain, (h5(e11) = h6(e12)), inference(forward_demodulation, [], [f4801, f1238])).
fof(f4801, plain, (op2(e24, e24) = h6(e12)), inference(forward_demodulation, [], [f1246, f1206])).
fof(f1246, plain, (op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))) = h6(e12)), inference(cnf_transformation, [], [f21])).
fof(f24106, plain, (~ (op2(e23, e22) = h6(e12)) | (~ spl42_178 | ~ spl42_743 | spl42_1040)), inference(forward_demodulation, [], [f24105, f2336])).
fof(f24105, plain, (~ (op2(e23, e22) = h6(op1(e13, e12))) | (~ spl42_743 | spl42_1040)), inference(forward_demodulation, [], [f6782, f5143])).
fof(f6782, plain, (~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | spl42_1040), inference(avatar_component_clause, [], [f6780])).
fof(f6780, plain, (spl42_1040 <=> (h6(op1(e13, e12)) = op2(e23, h5(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1040])])).
fof(f24058, plain, (~ spl42_155 | ~ spl42_498 | spl42_1037), inference(avatar_contradiction_clause, [], [f24057])).
fof(f24057, plain, ($false | (~ spl42_155 | ~ spl42_498 | spl42_1037)), inference(subsumption_resolution, [], [f24056, f3760])).
fof(f3760, plain, ((e20 = op2(e23, e25)) | ~ spl42_498), inference(avatar_component_clause, [], [f3758])).
fof(f3758, plain, (spl42_498 <=> (e20 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_498])])).
fof(f24056, plain, (~ (e20 = op2(e23, e25)) | (~ spl42_155 | spl42_1037)), inference(forward_demodulation, [], [f24055, f4820])).
fof(f24055, plain, (~ (op2(e23, e25) = h6(e10)) | (~ spl42_155 | spl42_1037)), inference(forward_demodulation, [], [f6770, f2241])).
fof(f2241, plain, ((e10 = op1(e13, e15)) | ~ spl42_155), inference(avatar_component_clause, [], [f2239])).
fof(f6770, plain, (~ (op2(e23, e25) = h6(op1(e13, e15))) | spl42_1037), inference(avatar_component_clause, [], [f6768])).
fof(f6768, plain, (spl42_1037 <=> (op2(e23, e25) = h6(op1(e13, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1037])])).
fof(f24009, plain, (~ spl42_126 | ~ spl42_469 | ~ spl42_1016 | spl42_1032), inference(avatar_contradiction_clause, [], [f24008])).
fof(f24008, plain, ($false | (~ spl42_126 | ~ spl42_469 | ~ spl42_1016 | spl42_1032)), inference(subsumption_resolution, [], [f24007, f3639])).
fof(f3639, plain, ((e26 = op2(e24, e23)) | ~ spl42_469), inference(avatar_component_clause, [], [f3637])).
fof(f3637, plain, (spl42_469 <=> (e26 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_469])])).
fof(f24007, plain, (~ (e26 = op2(e24, e23)) | (~ spl42_126 | ~ spl42_1016 | spl42_1032)), inference(forward_demodulation, [], [f24006, f6685])).
fof(f24006, plain, (~ (op2(e24, e23) = h6(e16)) | (~ spl42_126 | spl42_1032)), inference(forward_demodulation, [], [f6750, f2120])).
fof(f2120, plain, ((e16 = op1(e14, e13)) | ~ spl42_126), inference(avatar_component_clause, [], [f2118])).
fof(f2118, plain, (spl42_126 <=> (e16 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_126])])).
fof(f6750, plain, (~ (op2(e24, e23) = h6(op1(e14, e13))) | spl42_1032), inference(avatar_component_clause, [], [f6748])).
fof(f6748, plain, (spl42_1032 <=> (op2(e24, e23) = h6(op1(e14, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1032])])).
fof(f23956, plain, (~ spl42_100 | ~ spl42_443 | ~ spl42_748 | ~ spl42_1016 | spl42_1029), inference(avatar_contradiction_clause, [], [f23955])).
fof(f23955, plain, ($false | (~ spl42_100 | ~ spl42_443 | ~ spl42_748 | ~ spl42_1016 | spl42_1029)), inference(subsumption_resolution, [], [f23954, f3532])).
fof(f3532, plain, ((e21 = op2(e24, e26)) | ~ spl42_443), inference(avatar_component_clause, [], [f3530])).
fof(f3530, plain, (spl42_443 <=> (e21 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_443])])).
fof(f23954, plain, (~ (e21 = op2(e24, e26)) | (~ spl42_100 | ~ spl42_748 | ~ spl42_1016 | spl42_1029)), inference(forward_demodulation, [], [f23953, f5172])).
fof(f23953, plain, (~ (op2(e24, e26) = h6(e11)) | (~ spl42_100 | ~ spl42_1016 | spl42_1029)), inference(forward_demodulation, [], [f23952, f2013])).
fof(f2013, plain, ((e11 = op1(e14, e16)) | ~ spl42_100), inference(avatar_component_clause, [], [f2011])).
fof(f2011, plain, (spl42_100 <=> (e11 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_100])])).
fof(f23952, plain, (~ (op2(e24, e26) = h6(op1(e14, e16))) | (~ spl42_1016 | spl42_1029)), inference(forward_demodulation, [], [f6738, f6685])).
fof(f6738, plain, (~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | spl42_1029), inference(avatar_component_clause, [], [f6736])).
fof(f6736, plain, (spl42_1029 <=> (h6(op1(e14, e16)) = op2(e24, h6(e16)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1029])])).
fof(f23903, plain, (~ spl42_81 | ~ spl42_424 | ~ spl42_743 | spl42_1026), inference(avatar_contradiction_clause, [], [f23902])).
fof(f23902, plain, ($false | (~ spl42_81 | ~ spl42_424 | ~ spl42_743 | spl42_1026)), inference(subsumption_resolution, [], [f23901, f3453])).
fof(f3453, plain, ((e23 = op2(e25, e22)) | ~ spl42_424), inference(avatar_component_clause, [], [f3451])).
fof(f3451, plain, (spl42_424 <=> (e23 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_424])])).
fof(f23901, plain, (~ (e23 = op2(e25, e22)) | (~ spl42_81 | ~ spl42_743 | spl42_1026)), inference(forward_demodulation, [], [f23900, f4800])).
fof(f4800, plain, (e23 = h6(e13)), inference(forward_demodulation, [], [f4799, f4778])).
fof(f4778, plain, (e23 = op2(h5(e11), e25)), inference(backward_demodulation, [], [f4663, f1238])).
fof(f4663, plain, (e23 = op2(op2(e24, e24), e25)), inference(forward_demodulation, [], [f1205, f1206])).
fof(f1205, plain, (e23 = op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), e25)), inference(cnf_transformation, [], [f15])).
fof(f4799, plain, (h6(e13) = op2(h5(e11), e25)), inference(forward_demodulation, [], [f4798, f1238])).
fof(f4798, plain, (h6(e13) = op2(op2(e24, e24), e25)), inference(forward_demodulation, [], [f1247, f1206])).
fof(f1247, plain, (op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), e25) = h6(e13)), inference(cnf_transformation, [], [f21])).
fof(f23900, plain, (~ (op2(e25, e22) = h6(e13)) | (~ spl42_81 | ~ spl42_743 | spl42_1026)), inference(forward_demodulation, [], [f23899, f1934])).
fof(f1934, plain, ((e13 = op1(e15, e12)) | ~ spl42_81), inference(avatar_component_clause, [], [f1932])).
fof(f1932, plain, (spl42_81 <=> (e13 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_81])])).
fof(f23899, plain, (~ (op2(e25, e22) = h6(op1(e15, e12))) | (~ spl42_743 | spl42_1026)), inference(forward_demodulation, [], [f6726, f5143])).
fof(f6726, plain, (~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | spl42_1026), inference(avatar_component_clause, [], [f6724])).
fof(f6724, plain, (spl42_1026 <=> (h6(op1(e15, e12)) = op2(e25, h5(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1026])])).
fof(f23845, plain, (~ spl42_58 | spl42_1023), inference(avatar_contradiction_clause, [], [f23844])).
fof(f23844, plain, ($false | (~ spl42_58 | spl42_1023)), inference(trivial_inequality_removal, [], [f23843])).
fof(f23843, plain, (~ (h6(e11) = h6(e11)) | (~ spl42_58 | spl42_1023)), inference(forward_demodulation, [], [f6714, f1839])).
fof(f6714, plain, (~ (h6(e11) = h6(op1(e15, e15))) | spl42_1023), inference(avatar_component_clause, [], [f6712])).
fof(f6712, plain, (spl42_1023 <=> (h6(e11) = h6(op1(e15, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1023])])).
fof(f23780, plain, (~ spl42_329 | ~ spl42_672 | ~ spl42_743 | ~ spl42_1016 | spl42_1059), inference(avatar_contradiction_clause, [], [f23779])).
fof(f23779, plain, ($false | (~ spl42_329 | ~ spl42_672 | ~ spl42_743 | ~ spl42_1016 | spl42_1059)), inference(subsumption_resolution, [], [f23778, f4480])).
fof(f4480, plain, ((e26 = op2(e20, e22)) | ~ spl42_672), inference(avatar_component_clause, [], [f4478])).
fof(f4478, plain, (spl42_672 <=> (e26 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_672])])).
fof(f23778, plain, (~ (e26 = op2(e20, e22)) | (~ spl42_329 | ~ spl42_743 | ~ spl42_1016 | spl42_1059)), inference(forward_demodulation, [], [f23777, f6685])).
fof(f23777, plain, (~ (op2(e20, e22) = h6(e16)) | (~ spl42_329 | ~ spl42_743 | spl42_1059)), inference(forward_demodulation, [], [f23776, f2961])).
fof(f2961, plain, ((e16 = op1(e10, e12)) | ~ spl42_329), inference(avatar_component_clause, [], [f2959])).
fof(f2959, plain, (spl42_329 <=> (e16 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_329])])).
fof(f23776, plain, (~ (op2(e20, e22) = h6(op1(e10, e12))) | (~ spl42_743 | spl42_1059)), inference(forward_demodulation, [], [f6858, f5143])).
fof(f6858, plain, (~ (h6(op1(e10, e12)) = op2(e20, h5(e11))) | spl42_1059), inference(avatar_component_clause, [], [f6856])).
fof(f6856, plain, (spl42_1059 <=> (h6(op1(e10, e12)) = op2(e20, h5(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1059])])).
fof(f23719, plain, (~ spl42_304 | ~ spl42_647 | ~ spl42_743 | spl42_1056), inference(avatar_contradiction_clause, [], [f23718])).
fof(f23718, plain, ($false | (~ spl42_304 | ~ spl42_647 | ~ spl42_743 | spl42_1056)), inference(subsumption_resolution, [], [f23717, f4377])).
fof(f4377, plain, ((e22 = op2(e20, e25)) | ~ spl42_647), inference(avatar_component_clause, [], [f4375])).
fof(f4375, plain, (spl42_647 <=> (e22 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_647])])).
fof(f23717, plain, (~ (e22 = op2(e20, e25)) | (~ spl42_304 | ~ spl42_743 | spl42_1056)), inference(forward_demodulation, [], [f23716, f11296])).
fof(f23716, plain, (~ (op2(e20, e25) = h6(e12)) | (~ spl42_304 | spl42_1056)), inference(forward_demodulation, [], [f6846, f2858])).
fof(f2858, plain, ((e12 = op1(e10, e15)) | ~ spl42_304), inference(avatar_component_clause, [], [f2856])).
fof(f2856, plain, (spl42_304 <=> (e12 = op1(e10, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_304])])).
fof(f6846, plain, (~ (op2(e20, e25) = h6(op1(e10, e15))) | spl42_1056), inference(avatar_component_clause, [], [f6844])).
fof(f6844, plain, (spl42_1056 <=> (op2(e20, e25) = h6(op1(e10, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1056])])).
fof(f23659, plain, (~ spl42_263 | ~ spl42_606 | ~ spl42_748 | spl42_1051), inference(avatar_contradiction_clause, [], [f23658])).
fof(f23658, plain, ($false | (~ spl42_263 | ~ spl42_606 | ~ spl42_748 | spl42_1051)), inference(subsumption_resolution, [], [f23657, f4207])).
fof(f4207, plain, ((e23 = op2(e21, e24)) | ~ spl42_606), inference(avatar_component_clause, [], [f4205])).
fof(f4205, plain, (spl42_606 <=> (e23 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_606])])).
fof(f23657, plain, (~ (e23 = op2(e21, e24)) | (~ spl42_263 | ~ spl42_748 | spl42_1051)), inference(forward_demodulation, [], [f23656, f4800])).
fof(f23656, plain, (~ (op2(e21, e24) = h6(e13)) | (~ spl42_263 | ~ spl42_748 | spl42_1051)), inference(forward_demodulation, [], [f23655, f2688])).
fof(f2688, plain, ((e13 = op1(e11, e14)) | ~ spl42_263), inference(avatar_component_clause, [], [f2686])).
fof(f2686, plain, (spl42_263 <=> (e13 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_263])])).
fof(f23655, plain, (~ (op2(e21, e24) = h6(op1(e11, e14))) | (~ spl42_748 | spl42_1051)), inference(forward_demodulation, [], [f6826, f5172])).
fof(f6826, plain, (~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | spl42_1051), inference(avatar_component_clause, [], [f6824])).
fof(f6824, plain, (spl42_1051 <=> (h6(op1(e11, e14)) = op2(h6(e11), e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_1051])])).
fof(f23594, plain, (~ spl42_220 | ~ spl42_563 | ~ spl42_743 | spl42_1046), inference(avatar_contradiction_clause, [], [f23593])).
fof(f23593, plain, ($false | (~ spl42_220 | ~ spl42_563 | ~ spl42_743 | spl42_1046)), inference(subsumption_resolution, [], [f23592, f4029])).
fof(f4029, plain, ((e22 = op2(e22, e23)) | ~ spl42_563), inference(avatar_component_clause, [], [f4027])).
fof(f4027, plain, (spl42_563 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_563])])).
fof(f23592, plain, (~ (e22 = op2(e22, e23)) | (~ spl42_220 | ~ spl42_743 | spl42_1046)), inference(forward_demodulation, [], [f23591, f11296])).
fof(f23591, plain, (~ (op2(e22, e23) = h6(e12)) | (~ spl42_220 | ~ spl42_743 | spl42_1046)), inference(forward_demodulation, [], [f23590, f2510])).
fof(f2510, plain, ((e12 = op1(e12, e13)) | ~ spl42_220), inference(avatar_component_clause, [], [f2508])).
fof(f2508, plain, (spl42_220 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_220])])).
fof(f23590, plain, (~ (op2(e22, e23) = h6(op1(e12, e13))) | (~ spl42_743 | spl42_1046)), inference(forward_demodulation, [], [f6806, f5143])).
fof(f6806, plain, (~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | spl42_1046), inference(avatar_component_clause, [], [f6804])).
fof(f6804, plain, (spl42_1046 <=> (h6(op1(e12, e13)) = op2(h5(e11), e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_1046])])).
fof(f23530, plain, (~ spl42_201 | ~ spl42_544 | ~ spl42_743 | ~ spl42_1016 | spl42_1043), inference(avatar_contradiction_clause, [], [f23529])).
fof(f23529, plain, ($false | (~ spl42_201 | ~ spl42_544 | ~ spl42_743 | ~ spl42_1016 | spl42_1043)), inference(subsumption_resolution, [], [f23528, f3950])).
fof(f3950, plain, ((e24 = op2(e22, e26)) | ~ spl42_544), inference(avatar_component_clause, [], [f3948])).
fof(f3948, plain, (spl42_544 <=> (e24 = op2(e22, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_544])])).
fof(f23528, plain, (~ (e24 = op2(e22, e26)) | (~ spl42_201 | ~ spl42_743 | ~ spl42_1016 | spl42_1043)), inference(forward_demodulation, [], [f23527, f4797])).
fof(f4797, plain, (e24 = h6(e14)), inference(forward_demodulation, [], [f1248, f1206])).
fof(f1248, plain, (op2(e25, op2(e25, e25)) = h6(e14)), inference(cnf_transformation, [], [f21])).
fof(f23527, plain, (~ (op2(e22, e26) = h6(e14)) | (~ spl42_201 | ~ spl42_743 | ~ spl42_1016 | spl42_1043)), inference(forward_demodulation, [], [f23526, f2431])).
fof(f23526, plain, (~ (op2(e22, e26) = h6(op1(e12, e16))) | (~ spl42_743 | ~ spl42_1016 | spl42_1043)), inference(forward_demodulation, [], [f23525, f5143])).
fof(f23525, plain, (~ (h6(op1(e12, e16)) = op2(h5(e11), e26)) | (~ spl42_1016 | spl42_1043)), inference(forward_demodulation, [], [f6794, f6685])).
fof(f6794, plain, (~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | spl42_1043), inference(avatar_component_clause, [], [f6792])).
fof(f6792, plain, (spl42_1043 <=> (h6(op1(e12, e16)) = op2(h5(e11), h6(e16)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1043])])).
fof(f23463, plain, (~ spl42_173 | ~ spl42_801 | spl42_1039), inference(avatar_contradiction_clause, [], [f23462])).
fof(f23462, plain, ($false | (~ spl42_173 | ~ spl42_801 | spl42_1039)), inference(subsumption_resolution, [], [f23461, f5450])).
fof(f5450, plain, ((e24 = h4(e11)) | ~ spl42_801), inference(avatar_component_clause, [], [f5449])).
fof(f5449, plain, (spl42_801 <=> (e24 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_801])])).
fof(f23461, plain, (~ (e24 = h4(e11)) | (~ spl42_173 | spl42_1039)), inference(forward_demodulation, [], [f23460, f4797])).
fof(f23460, plain, (~ (h4(e11) = h6(e14)) | (~ spl42_173 | spl42_1039)), inference(forward_demodulation, [], [f6778, f2315])).
fof(f2315, plain, ((e14 = op1(e13, e13)) | ~ spl42_173), inference(avatar_component_clause, [], [f2313])).
fof(f2313, plain, (spl42_173 <=> (e14 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_173])])).
fof(f6778, plain, (~ (h4(e11) = h6(op1(e13, e13))) | spl42_1039), inference(avatar_component_clause, [], [f6776])).
fof(f6776, plain, (spl42_1039 <=> (h4(e11) = h6(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1039])])).
fof(f23400, plain, (~ spl42_137 | ~ spl42_480 | ~ spl42_748 | spl42_1034), inference(avatar_contradiction_clause, [], [f23399])).
fof(f23399, plain, ($false | (~ spl42_137 | ~ spl42_480 | ~ spl42_748 | spl42_1034)), inference(subsumption_resolution, [], [f23398, f3685])).
fof(f3685, plain, ((e23 = op2(e24, e21)) | ~ spl42_480), inference(avatar_component_clause, [], [f3683])).
fof(f3683, plain, (spl42_480 <=> (e23 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_480])])).
fof(f23398, plain, (~ (e23 = op2(e24, e21)) | (~ spl42_137 | ~ spl42_748 | spl42_1034)), inference(forward_demodulation, [], [f23397, f4800])).
fof(f23397, plain, (~ (op2(e24, e21) = h6(e13)) | (~ spl42_137 | ~ spl42_748 | spl42_1034)), inference(forward_demodulation, [], [f23396, f2166])).
fof(f2166, plain, ((e13 = op1(e14, e11)) | ~ spl42_137), inference(avatar_component_clause, [], [f2164])).
fof(f2164, plain, (spl42_137 <=> (e13 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_137])])).
fof(f23396, plain, (~ (op2(e24, e21) = h6(op1(e14, e11))) | (~ spl42_748 | spl42_1034)), inference(forward_demodulation, [], [f6758, f5172])).
fof(f6758, plain, (~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | spl42_1034), inference(avatar_component_clause, [], [f6756])).
fof(f6756, plain, (spl42_1034 <=> (h6(op1(e14, e11)) = op2(e24, h6(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1034])])).
fof(f23338, plain, (~ spl42_111 | ~ spl42_454 | spl42_1030), inference(avatar_contradiction_clause, [], [f23337])).
fof(f23337, plain, ($false | (~ spl42_111 | ~ spl42_454 | spl42_1030)), inference(subsumption_resolution, [], [f23336, f3577])).
fof(f3577, plain, ((e25 = op2(e24, e25)) | ~ spl42_454), inference(avatar_component_clause, [], [f3575])).
fof(f3575, plain, (spl42_454 <=> (e25 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_454])])).
fof(f23336, plain, (~ (e25 = op2(e24, e25)) | (~ spl42_111 | spl42_1030)), inference(forward_demodulation, [], [f23335, f1243])).
fof(f1243, plain, (e25 = h6(e15)), inference(cnf_transformation, [], [f21])).
fof(f23335, plain, (~ (op2(e24, e25) = h6(e15)) | (~ spl42_111 | spl42_1030)), inference(forward_demodulation, [], [f6742, f2058])).
fof(f2058, plain, ((e15 = op1(e14, e15)) | ~ spl42_111), inference(avatar_component_clause, [], [f2056])).
fof(f2056, plain, (spl42_111 <=> (e15 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_111])])).
fof(f6742, plain, (~ (op2(e24, e25) = h6(op1(e14, e15))) | spl42_1030), inference(avatar_component_clause, [], [f6740])).
fof(f6740, plain, (spl42_1030 <=> (op2(e24, e25) = h6(op1(e14, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1030])])).
fof(f23276, plain, (~ spl42_94 | ~ spl42_437 | ~ spl42_743 | spl42_1028), inference(avatar_contradiction_clause, [], [f23275])).
fof(f23275, plain, ($false | (~ spl42_94 | ~ spl42_437 | ~ spl42_743 | spl42_1028)), inference(subsumption_resolution, [], [f23274, f3507])).
fof(f3507, plain, ((e22 = op2(e25, e20)) | ~ spl42_437), inference(avatar_component_clause, [], [f3505])).
fof(f3505, plain, (spl42_437 <=> (e22 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_437])])).
fof(f23274, plain, (~ (e22 = op2(e25, e20)) | (~ spl42_94 | ~ spl42_743 | spl42_1028)), inference(forward_demodulation, [], [f23273, f11296])).
fof(f23273, plain, (~ (op2(e25, e20) = h6(e12)) | (~ spl42_94 | spl42_1028)), inference(forward_demodulation, [], [f6734, f1988])).
fof(f1988, plain, ((e12 = op1(e15, e10)) | ~ spl42_94), inference(avatar_component_clause, [], [f1986])).
fof(f1986, plain, (spl42_94 <=> (e12 = op1(e15, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_94])])).
fof(f6734, plain, (~ (op2(e25, e20) = h6(op1(e15, e10))) | spl42_1028), inference(avatar_component_clause, [], [f6732])).
fof(f6732, plain, (spl42_1028 <=> (op2(e25, e20) = h6(op1(e15, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1028])])).
fof(f23215, plain, (~ spl42_69 | ~ spl42_412 | spl42_1024), inference(avatar_contradiction_clause, [], [f23214])).
fof(f23214, plain, ($false | (~ spl42_69 | ~ spl42_412 | spl42_1024)), inference(subsumption_resolution, [], [f23213, f3403])).
fof(f3403, plain, ((e25 = op2(e25, e24)) | ~ spl42_412), inference(avatar_component_clause, [], [f3401])).
fof(f3401, plain, (spl42_412 <=> (e25 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_412])])).
fof(f23213, plain, (~ (e25 = op2(e25, e24)) | (~ spl42_69 | spl42_1024)), inference(forward_demodulation, [], [f23212, f1243])).
fof(f23212, plain, (~ (op2(e25, e24) = h6(e15)) | (~ spl42_69 | spl42_1024)), inference(forward_demodulation, [], [f6718, f1884])).
fof(f1884, plain, ((e15 = op1(e15, e14)) | ~ spl42_69), inference(avatar_component_clause, [], [f1882])).
fof(f1882, plain, (spl42_69 <=> (e15 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_69])])).
fof(f6718, plain, (~ (op2(e25, e24) = h6(op1(e15, e14))) | spl42_1024), inference(avatar_component_clause, [], [f6716])).
fof(f6716, plain, (spl42_1024 <=> (op2(e25, e24) = h6(op1(e15, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1024])])).
fof(f23153, plain, (~ spl42_56 | ~ spl42_399 | ~ spl42_1016 | spl42_1022), inference(avatar_contradiction_clause, [], [f23152])).
fof(f23152, plain, ($false | (~ spl42_56 | ~ spl42_399 | ~ spl42_1016 | spl42_1022)), inference(subsumption_resolution, [], [f23151, f3349])).
fof(f3349, plain, ((e26 = op2(e25, e26)) | ~ spl42_399), inference(avatar_component_clause, [], [f3347])).
fof(f3347, plain, (spl42_399 <=> (e26 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_399])])).
fof(f23151, plain, (~ (e26 = op2(e25, e26)) | (~ spl42_56 | ~ spl42_1016 | spl42_1022)), inference(forward_demodulation, [], [f23150, f6685])).
fof(f23150, plain, (~ (op2(e25, e26) = h6(e16)) | (~ spl42_56 | ~ spl42_1016 | spl42_1022)), inference(forward_demodulation, [], [f23149, f1830])).
fof(f1830, plain, ((e16 = op1(e15, e16)) | ~ spl42_56), inference(avatar_component_clause, [], [f1828])).
fof(f1828, plain, (spl42_56 <=> (e16 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_56])])).
fof(f23149, plain, (~ (op2(e25, e26) = h6(op1(e15, e16))) | (~ spl42_1016 | spl42_1022)), inference(forward_demodulation, [], [f6710, f6685])).
fof(f6710, plain, (~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | spl42_1022), inference(avatar_component_clause, [], [f6708])).
fof(f6708, plain, (spl42_1022 <=> (h6(op1(e15, e16)) = op2(e25, h6(e16)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1022])])).
fof(f23086, plain, (~ spl42_48 | ~ spl42_391 | ~ spl42_1016 | spl42_1021), inference(avatar_contradiction_clause, [], [f23085])).
fof(f23085, plain, ($false | (~ spl42_48 | ~ spl42_391 | ~ spl42_1016 | spl42_1021)), inference(subsumption_resolution, [], [f23084, f3316])).
fof(f3316, plain, ((e25 = op2(e26, e20)) | ~ spl42_391), inference(avatar_component_clause, [], [f3314])).
fof(f3314, plain, (spl42_391 <=> (e25 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_391])])).
fof(f23084, plain, (~ (e25 = op2(e26, e20)) | (~ spl42_48 | ~ spl42_1016 | spl42_1021)), inference(forward_demodulation, [], [f23083, f1243])).
fof(f23083, plain, (~ (op2(e26, e20) = h6(e15)) | (~ spl42_48 | ~ spl42_1016 | spl42_1021)), inference(forward_demodulation, [], [f23082, f1797])).
fof(f23082, plain, (~ (op2(e26, e20) = h6(op1(e16, e10))) | (~ spl42_1016 | spl42_1021)), inference(forward_demodulation, [], [f6706, f6685])).
fof(f6706, plain, (~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | spl42_1021), inference(avatar_component_clause, [], [f6704])).
fof(f6704, plain, (spl42_1021 <=> (h6(op1(e16, e10)) = op2(h6(e16), e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_1021])])).
fof(f23020, plain, (~ spl42_33 | ~ spl42_376 | ~ spl42_743 | ~ spl42_1016 | spl42_1020), inference(avatar_contradiction_clause, [], [f23019])).
fof(f23019, plain, ($false | (~ spl42_33 | ~ spl42_376 | ~ spl42_743 | ~ spl42_1016 | spl42_1020)), inference(subsumption_resolution, [], [f23018, f3254])).
fof(f3254, plain, ((e24 = op2(e26, e22)) | ~ spl42_376), inference(avatar_component_clause, [], [f3252])).
fof(f3252, plain, (spl42_376 <=> (e24 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_376])])).
fof(f23018, plain, (~ (e24 = op2(e26, e22)) | (~ spl42_33 | ~ spl42_743 | ~ spl42_1016 | spl42_1020)), inference(forward_demodulation, [], [f23017, f4797])).
fof(f23017, plain, (~ (op2(e26, e22) = h6(e14)) | (~ spl42_33 | ~ spl42_743 | ~ spl42_1016 | spl42_1020)), inference(forward_demodulation, [], [f23016, f1735])).
fof(f1735, plain, ((e14 = op1(e16, e12)) | ~ spl42_33), inference(avatar_component_clause, [], [f1733])).
fof(f1733, plain, (spl42_33 <=> (e14 = op1(e16, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_33])])).
fof(f23016, plain, (~ (op2(e26, e22) = h6(op1(e16, e12))) | (~ spl42_743 | ~ spl42_1016 | spl42_1020)), inference(forward_demodulation, [], [f23015, f6685])).
fof(f23015, plain, (~ (h6(op1(e16, e12)) = op2(h6(e16), e22)) | (~ spl42_743 | spl42_1020)), inference(forward_demodulation, [], [f6702, f5143])).
fof(f6702, plain, (~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | spl42_1020), inference(avatar_component_clause, [], [f6700])).
fof(f6700, plain, (spl42_1020 <=> (h6(op1(e16, e12)) = op2(h6(e16), h5(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1020])])).
fof(f22954, plain, (~ spl42_25 | ~ spl42_368 | ~ spl42_1016 | spl42_1019), inference(avatar_contradiction_clause, [], [f22953])).
fof(f22953, plain, ($false | (~ spl42_25 | ~ spl42_368 | ~ spl42_1016 | spl42_1019)), inference(subsumption_resolution, [], [f22952, f3221])).
fof(f3221, plain, ((e23 = op2(e26, e23)) | ~ spl42_368), inference(avatar_component_clause, [], [f3219])).
fof(f3219, plain, (spl42_368 <=> (e23 = op2(e26, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_368])])).
fof(f22952, plain, (~ (e23 = op2(e26, e23)) | (~ spl42_25 | ~ spl42_1016 | spl42_1019)), inference(forward_demodulation, [], [f22951, f4800])).
fof(f22951, plain, (~ (op2(e26, e23) = h6(e13)) | (~ spl42_25 | ~ spl42_1016 | spl42_1019)), inference(forward_demodulation, [], [f22950, f1702])).
fof(f1702, plain, ((e13 = op1(e16, e13)) | ~ spl42_25), inference(avatar_component_clause, [], [f1700])).
fof(f1700, plain, (spl42_25 <=> (e13 = op1(e16, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_25])])).
fof(f22950, plain, (~ (op2(e26, e23) = h6(op1(e16, e13))) | (~ spl42_1016 | spl42_1019)), inference(forward_demodulation, [], [f6698, f6685])).
fof(f6698, plain, (~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | spl42_1019), inference(avatar_component_clause, [], [f6696])).
fof(f6696, plain, (spl42_1019 <=> (h6(op1(e16, e13)) = op2(h6(e16), e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_1019])])).
fof(f22893, plain, (~ spl42_38 | ~ spl42_381 | ~ spl42_743 | ~ spl42_748 | spl42_1014 | ~ spl42_1016), inference(avatar_contradiction_clause, [], [f22892])).
fof(f22892, plain, ($false | (~ spl42_38 | ~ spl42_381 | ~ spl42_743 | ~ spl42_748 | spl42_1014 | ~ spl42_1016)), inference(subsumption_resolution, [], [f22891, f3275])).
fof(f3275, plain, ((e22 = op2(e26, e21)) | ~ spl42_381), inference(avatar_component_clause, [], [f3273])).
fof(f3273, plain, (spl42_381 <=> (e22 = op2(e26, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_381])])).
fof(f22891, plain, (~ (e22 = op2(e26, e21)) | (~ spl42_38 | ~ spl42_743 | ~ spl42_748 | spl42_1014 | ~ spl42_1016)), inference(forward_demodulation, [], [f22890, f11296])).
fof(f22890, plain, (~ (op2(e26, e21) = h6(e12)) | (~ spl42_38 | ~ spl42_748 | spl42_1014 | ~ spl42_1016)), inference(forward_demodulation, [], [f22889, f1756])).
fof(f1756, plain, ((e12 = op1(e16, e11)) | ~ spl42_38), inference(avatar_component_clause, [], [f1754])).
fof(f1754, plain, (spl42_38 <=> (e12 = op1(e16, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_38])])).
fof(f22889, plain, (~ (op2(e26, e21) = h6(op1(e16, e11))) | (~ spl42_748 | spl42_1014 | ~ spl42_1016)), inference(forward_demodulation, [], [f22888, f6685])).
fof(f22888, plain, (~ (h6(op1(e16, e11)) = op2(h6(e16), e21)) | (~ spl42_748 | spl42_1014)), inference(forward_demodulation, [], [f6678, f5172])).
fof(f6678, plain, (~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | spl42_1014), inference(avatar_component_clause, [], [f6676])).
fof(f6676, plain, (spl42_1014 <=> (h6(op1(e16, e11)) = op2(h6(e16), h6(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1014])])).
fof(f22832, plain, (~ spl42_248 | ~ spl42_591 | ~ spl42_743 | ~ spl42_748 | spl42_1013 | ~ spl42_1016), inference(avatar_contradiction_clause, [], [f22831])).
fof(f22831, plain, ($false | (~ spl42_248 | ~ spl42_591 | ~ spl42_743 | ~ spl42_748 | spl42_1013 | ~ spl42_1016)), inference(subsumption_resolution, [], [f22830, f4145])).
fof(f4145, plain, ((e22 = op2(e21, e26)) | ~ spl42_591), inference(avatar_component_clause, [], [f4143])).
fof(f4143, plain, (spl42_591 <=> (e22 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_591])])).
fof(f22830, plain, (~ (e22 = op2(e21, e26)) | (~ spl42_248 | ~ spl42_743 | ~ spl42_748 | spl42_1013 | ~ spl42_1016)), inference(forward_demodulation, [], [f22829, f11296])).
fof(f22829, plain, (~ (op2(e21, e26) = h6(e12)) | (~ spl42_248 | ~ spl42_748 | spl42_1013 | ~ spl42_1016)), inference(forward_demodulation, [], [f22828, f2626])).
fof(f2626, plain, ((e12 = op1(e11, e16)) | ~ spl42_248), inference(avatar_component_clause, [], [f2624])).
fof(f2624, plain, (spl42_248 <=> (e12 = op1(e11, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_248])])).
fof(f22828, plain, (~ (op2(e21, e26) = h6(op1(e11, e16))) | (~ spl42_748 | spl42_1013 | ~ spl42_1016)), inference(forward_demodulation, [], [f22827, f5172])).
fof(f22827, plain, (~ (h6(op1(e11, e16)) = op2(h6(e11), e26)) | (spl42_1013 | ~ spl42_1016)), inference(forward_demodulation, [], [f6674, f6685])).
fof(f6674, plain, (~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | spl42_1013), inference(avatar_component_clause, [], [f6672])).
fof(f6672, plain, (spl42_1013 <=> (h6(op1(e11, e16)) = op2(h6(e11), h6(e16)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1013])])).
fof(f22767, plain, (~ spl42_340 | ~ spl42_934 | spl42_1061), inference(avatar_contradiction_clause, [], [f22766])).
fof(f22766, plain, ($false | (~ spl42_340 | ~ spl42_934 | spl42_1061)), inference(subsumption_resolution, [], [f22765, f6118])).
fof(f6118, plain, ((e23 = h1(e11)) | ~ spl42_934), inference(avatar_component_clause, [], [f6117])).
fof(f6117, plain, (spl42_934 <=> (e23 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_934])])).
fof(f22765, plain, (~ (e23 = h1(e11)) | (~ spl42_340 | spl42_1061)), inference(forward_demodulation, [], [f22764, f4800])).
fof(f22764, plain, (~ (h1(e11) = h6(e13)) | (~ spl42_340 | spl42_1061)), inference(forward_demodulation, [], [f6866, f3007])).
fof(f3007, plain, ((op1(e10, e10) = e13) | ~ spl42_340), inference(avatar_component_clause, [], [f3005])).
fof(f3005, plain, (spl42_340 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl42_340])])).
fof(f6866, plain, (~ (h1(e11) = h6(op1(e10, e10))) | spl42_1061), inference(avatar_component_clause, [], [f6864])).
fof(f6864, plain, (spl42_1061 <=> (h1(e11) = h6(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1061])])).
fof(f22707, plain, (~ spl42_330 | ~ spl42_673 | ~ spl42_748 | spl42_1060), inference(avatar_contradiction_clause, [], [f22706])).
fof(f22706, plain, ($false | (~ spl42_330 | ~ spl42_673 | ~ spl42_748 | spl42_1060)), inference(subsumption_resolution, [], [f22705, f4485])).
fof(f4485, plain, ((e20 = op2(e20, e21)) | ~ spl42_673), inference(avatar_component_clause, [], [f4483])).
fof(f4483, plain, (spl42_673 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_673])])).
fof(f22705, plain, (~ (e20 = op2(e20, e21)) | (~ spl42_330 | ~ spl42_748 | spl42_1060)), inference(forward_demodulation, [], [f22704, f4820])).
fof(f22704, plain, (~ (op2(e20, e21) = h6(e10)) | (~ spl42_330 | ~ spl42_748 | spl42_1060)), inference(forward_demodulation, [], [f22703, f2966])).
fof(f2966, plain, ((e10 = op1(e10, e11)) | ~ spl42_330), inference(avatar_component_clause, [], [f2964])).
fof(f2964, plain, (spl42_330 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_330])])).
fof(f22703, plain, (~ (op2(e20, e21) = h6(op1(e10, e11))) | (~ spl42_748 | spl42_1060)), inference(forward_demodulation, [], [f6862, f5172])).
fof(f6862, plain, (~ (h6(op1(e10, e11)) = op2(e20, h6(e11))) | spl42_1060), inference(avatar_component_clause, [], [f6860])).
fof(f6860, plain, (spl42_1060 <=> (h6(op1(e10, e11)) = op2(e20, h6(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1060])])).
fof(f22642, plain, (~ spl42_313 | ~ spl42_656 | spl42_1057), inference(avatar_contradiction_clause, [], [f22641])).
fof(f22641, plain, ($false | (~ spl42_313 | ~ spl42_656 | spl42_1057)), inference(subsumption_resolution, [], [f22640, f4414])).
fof(f4414, plain, ((e24 = op2(e20, e24)) | ~ spl42_656), inference(avatar_component_clause, [], [f4412])).
fof(f4412, plain, (spl42_656 <=> (e24 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_656])])).
fof(f22640, plain, (~ (e24 = op2(e20, e24)) | (~ spl42_313 | spl42_1057)), inference(forward_demodulation, [], [f22639, f4797])).
fof(f22639, plain, (~ (op2(e20, e24) = h6(e14)) | (~ spl42_313 | spl42_1057)), inference(forward_demodulation, [], [f6850, f2895])).
fof(f2895, plain, ((e14 = op1(e10, e14)) | ~ spl42_313), inference(avatar_component_clause, [], [f2893])).
fof(f2893, plain, (spl42_313 <=> (e14 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_313])])).
fof(f6850, plain, (~ (op2(e20, e24) = h6(op1(e10, e14))) | spl42_1057), inference(avatar_component_clause, [], [f6848])).
fof(f6848, plain, (spl42_1057 <=> (op2(e20, e24) = h6(op1(e10, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1057])])).
fof(f22579, plain, (~ spl42_300 | ~ spl42_643 | ~ spl42_1016 | spl42_1055), inference(avatar_contradiction_clause, [], [f22578])).
fof(f22578, plain, ($false | (~ spl42_300 | ~ spl42_643 | ~ spl42_1016 | spl42_1055)), inference(subsumption_resolution, [], [f22577, f4360])).
fof(f4360, plain, ((e25 = op2(e20, e26)) | ~ spl42_643), inference(avatar_component_clause, [], [f4358])).
fof(f4358, plain, (spl42_643 <=> (e25 = op2(e20, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_643])])).
fof(f22577, plain, (~ (e25 = op2(e20, e26)) | (~ spl42_300 | ~ spl42_1016 | spl42_1055)), inference(forward_demodulation, [], [f22576, f1243])).
fof(f22576, plain, (~ (op2(e20, e26) = h6(e15)) | (~ spl42_300 | ~ spl42_1016 | spl42_1055)), inference(forward_demodulation, [], [f22575, f2841])).
fof(f2841, plain, ((e15 = op1(e10, e16)) | ~ spl42_300), inference(avatar_component_clause, [], [f2839])).
fof(f2839, plain, (spl42_300 <=> (e15 = op1(e10, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_300])])).
fof(f22575, plain, (~ (op2(e20, e26) = h6(op1(e10, e16))) | (~ spl42_1016 | spl42_1055)), inference(forward_demodulation, [], [f6842, f6685])).
fof(f6842, plain, (~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | spl42_1055), inference(avatar_component_clause, [], [f6840])).
fof(f6840, plain, (spl42_1055 <=> (h6(op1(e10, e16)) = op2(e20, h6(e16)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1055])])).
fof(f22520, plain, (~ spl42_272 | ~ spl42_615 | ~ spl42_748 | spl42_1052), inference(avatar_contradiction_clause, [], [f22519])).
fof(f22519, plain, ($false | (~ spl42_272 | ~ spl42_615 | ~ spl42_748 | spl42_1052)), inference(subsumption_resolution, [], [f22518, f4244])).
fof(f4244, plain, ((e25 = op2(e21, e23)) | ~ spl42_615), inference(avatar_component_clause, [], [f4242])).
fof(f4242, plain, (spl42_615 <=> (e25 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_615])])).
fof(f22518, plain, (~ (e25 = op2(e21, e23)) | (~ spl42_272 | ~ spl42_748 | spl42_1052)), inference(forward_demodulation, [], [f22517, f1243])).
fof(f22517, plain, (~ (op2(e21, e23) = h6(e15)) | (~ spl42_272 | ~ spl42_748 | spl42_1052)), inference(forward_demodulation, [], [f22516, f2725])).
fof(f2725, plain, ((e15 = op1(e11, e13)) | ~ spl42_272), inference(avatar_component_clause, [], [f2723])).
fof(f2723, plain, (spl42_272 <=> (e15 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_272])])).
fof(f22516, plain, (~ (op2(e21, e23) = h6(op1(e11, e13))) | (~ spl42_748 | spl42_1052)), inference(forward_demodulation, [], [f6830, f5172])).
fof(f6830, plain, (~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | spl42_1052), inference(avatar_component_clause, [], [f6828])).
fof(f6828, plain, (spl42_1052 <=> (h6(op1(e11, e13)) = op2(h6(e11), e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_1052])])).
fof(f22459, plain, (~ spl42_257 | ~ spl42_600 | ~ spl42_748 | spl42_1050), inference(avatar_contradiction_clause, [], [f22458])).
fof(f22458, plain, ($false | (~ spl42_257 | ~ spl42_600 | ~ spl42_748 | spl42_1050)), inference(subsumption_resolution, [], [f22457, f4182])).
fof(f4182, plain, ((e24 = op2(e21, e25)) | ~ spl42_600), inference(avatar_component_clause, [], [f4180])).
fof(f4180, plain, (spl42_600 <=> (e24 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_600])])).
fof(f22457, plain, (~ (e24 = op2(e21, e25)) | (~ spl42_257 | ~ spl42_748 | spl42_1050)), inference(forward_demodulation, [], [f22456, f4797])).
fof(f22456, plain, (~ (op2(e21, e25) = h6(e14)) | (~ spl42_257 | ~ spl42_748 | spl42_1050)), inference(forward_demodulation, [], [f22455, f2663])).
fof(f2663, plain, ((e14 = op1(e11, e15)) | ~ spl42_257), inference(avatar_component_clause, [], [f2661])).
fof(f22455, plain, (~ (op2(e21, e25) = h6(op1(e11, e15))) | (~ spl42_748 | spl42_1050)), inference(forward_demodulation, [], [f6822, f5172])).
fof(f6822, plain, (~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | spl42_1050), inference(avatar_component_clause, [], [f6820])).
fof(f6820, plain, (spl42_1050 <=> (h6(op1(e11, e15)) = op2(h6(e11), e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_1050])])).
fof(f22394, plain, (~ spl42_233 | ~ spl42_576 | ~ spl42_743 | ~ spl42_748 | spl42_1048), inference(avatar_contradiction_clause, [], [f22393])).
fof(f22393, plain, ($false | (~ spl42_233 | ~ spl42_576 | ~ spl42_743 | ~ spl42_748 | spl42_1048)), inference(subsumption_resolution, [], [f22392, f4083])).
fof(f4083, plain, ((e21 = op2(e22, e21)) | ~ spl42_576), inference(avatar_component_clause, [], [f4081])).
fof(f4081, plain, (spl42_576 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_576])])).
fof(f22392, plain, (~ (e21 = op2(e22, e21)) | (~ spl42_233 | ~ spl42_743 | ~ spl42_748 | spl42_1048)), inference(forward_demodulation, [], [f22391, f5172])).
fof(f22391, plain, (~ (op2(e22, e21) = h6(e11)) | (~ spl42_233 | ~ spl42_743 | ~ spl42_748 | spl42_1048)), inference(forward_demodulation, [], [f22390, f2564])).
fof(f2564, plain, ((e11 = op1(e12, e11)) | ~ spl42_233), inference(avatar_component_clause, [], [f2562])).
fof(f2562, plain, (spl42_233 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_233])])).
fof(f22390, plain, (~ (op2(e22, e21) = h6(op1(e12, e11))) | (~ spl42_743 | ~ spl42_748 | spl42_1048)), inference(forward_demodulation, [], [f22389, f5143])).
fof(f22389, plain, (~ (h6(op1(e12, e11)) = op2(h5(e11), e21)) | (~ spl42_748 | spl42_1048)), inference(forward_demodulation, [], [f6814, f5172])).
fof(f6814, plain, (~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | spl42_1048), inference(avatar_component_clause, [], [f6812])).
fof(f6812, plain, (spl42_1048 <=> (h6(op1(e12, e11)) = op2(h5(e11), h6(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1048])])).
fof(f22329, plain, (~ spl42_230 | ~ spl42_573 | ~ spl42_743 | spl42_1047), inference(avatar_contradiction_clause, [], [f22328])).
fof(f22328, plain, ($false | (~ spl42_230 | ~ spl42_573 | ~ spl42_743 | spl42_1047)), inference(subsumption_resolution, [], [f22327, f4070])).
fof(f4070, plain, ((e25 = op2(e22, e22)) | ~ spl42_573), inference(avatar_component_clause, [], [f4068])).
fof(f4068, plain, (spl42_573 <=> (e25 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_573])])).
fof(f22327, plain, (~ (e25 = op2(e22, e22)) | (~ spl42_230 | ~ spl42_743 | spl42_1047)), inference(forward_demodulation, [], [f22326, f1243])).
fof(f22326, plain, (~ (op2(e22, e22) = h6(e15)) | (~ spl42_230 | ~ spl42_743 | spl42_1047)), inference(forward_demodulation, [], [f22325, f5143])).
fof(f22325, plain, (~ (h6(e15) = op2(h5(e11), h5(e11))) | (~ spl42_230 | spl42_1047)), inference(forward_demodulation, [], [f6810, f2551])).
fof(f6810, plain, (~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | spl42_1047), inference(avatar_component_clause, [], [f6808])).
fof(f6808, plain, (spl42_1047 <=> (op2(h5(e11), h5(e11)) = h6(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1047])])).
fof(f22268, plain, (~ spl42_207 | spl42_1044), inference(avatar_contradiction_clause, [], [f22267])).
fof(f22267, plain, ($false | (~ spl42_207 | spl42_1044)), inference(subsumption_resolution, [], [f22266, f4800])).
fof(f22266, plain, (~ (e23 = h6(e13)) | (~ spl42_207 | spl42_1044)), inference(forward_demodulation, [], [f6798, f2456])).
fof(f2456, plain, ((e13 = op1(e12, e15)) | ~ spl42_207), inference(avatar_component_clause, [], [f2454])).
fof(f2454, plain, (spl42_207 <=> (e13 = op1(e12, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_207])])).
fof(f6798, plain, (~ (e23 = h6(op1(e12, e15))) | spl42_1044), inference(avatar_component_clause, [], [f6796])).
fof(f6796, plain, (spl42_1044 <=> (e23 = h6(op1(e12, e15)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1044])])).
fof(f22207, plain, (~ spl42_191 | ~ spl42_534 | ~ spl42_748 | spl42_1042), inference(avatar_contradiction_clause, [], [f22206])).
fof(f22206, plain, ($false | (~ spl42_191 | ~ spl42_534 | ~ spl42_748 | spl42_1042)), inference(subsumption_resolution, [], [f22205, f3909])).
fof(f3909, plain, ((e21 = op2(e23, e20)) | ~ spl42_534), inference(avatar_component_clause, [], [f3907])).
fof(f3907, plain, (spl42_534 <=> (e21 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_534])])).
fof(f22205, plain, (~ (e21 = op2(e23, e20)) | (~ spl42_191 | ~ spl42_748 | spl42_1042)), inference(forward_demodulation, [], [f22204, f5172])).
fof(f22204, plain, (~ (op2(e23, e20) = h6(e11)) | (~ spl42_191 | spl42_1042)), inference(forward_demodulation, [], [f6790, f2390])).
fof(f2390, plain, ((e11 = op1(e13, e10)) | ~ spl42_191), inference(avatar_component_clause, [], [f2388])).
fof(f2388, plain, (spl42_191 <=> (e11 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_191])])).
fof(f6790, plain, (~ (op2(e23, e20) = h6(op1(e13, e10))) | spl42_1042), inference(avatar_component_clause, [], [f6788])).
fof(f6788, plain, (spl42_1042 <=> (op2(e23, e20) = h6(op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1042])])).
fof(f22146, plain, (~ spl42_188 | ~ spl42_531 | ~ spl42_748 | spl42_1041), inference(avatar_contradiction_clause, [], [f22145])).
fof(f22145, plain, ($false | (~ spl42_188 | ~ spl42_531 | ~ spl42_748 | spl42_1041)), inference(subsumption_resolution, [], [f22144, f3896])).
fof(f3896, plain, ((e25 = op2(e23, e21)) | ~ spl42_531), inference(avatar_component_clause, [], [f3894])).
fof(f3894, plain, (spl42_531 <=> (e25 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_531])])).
fof(f22144, plain, (~ (e25 = op2(e23, e21)) | (~ spl42_188 | ~ spl42_748 | spl42_1041)), inference(forward_demodulation, [], [f22143, f1243])).
fof(f22143, plain, (~ (op2(e23, e21) = h6(e15)) | (~ spl42_188 | ~ spl42_748 | spl42_1041)), inference(forward_demodulation, [], [f22142, f2377])).
fof(f2377, plain, ((e15 = op1(e13, e11)) | ~ spl42_188), inference(avatar_component_clause, [], [f2375])).
fof(f2375, plain, (spl42_188 <=> (e15 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_188])])).
fof(f22142, plain, (~ (op2(e23, e21) = h6(op1(e13, e11))) | (~ spl42_748 | spl42_1041)), inference(forward_demodulation, [], [f6786, f5172])).
fof(f6786, plain, (~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | spl42_1041), inference(avatar_component_clause, [], [f6784])).
fof(f6784, plain, (spl42_1041 <=> (h6(op1(e13, e11)) = op2(e23, h6(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1041])])).
fof(f22085, plain, (~ spl42_168 | spl42_1038), inference(avatar_contradiction_clause, [], [f22084])).
fof(f22084, plain, ($false | (~ spl42_168 | spl42_1038)), inference(trivial_inequality_removal, [], [f22083])).
fof(f22083, plain, (~ (h6(e16) = h6(e16)) | (~ spl42_168 | spl42_1038)), inference(forward_demodulation, [], [f6774, f2294])).
fof(f2294, plain, ((e16 = op1(e13, e14)) | ~ spl42_168), inference(avatar_component_clause, [], [f2292])).
fof(f2292, plain, (spl42_168 <=> (e16 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_168])])).
fof(f6774, plain, (~ (h6(e16) = h6(op1(e13, e14))) | spl42_1038), inference(avatar_component_clause, [], [f6772])).
fof(f6772, plain, (spl42_1038 <=> (h6(e16) = h6(op1(e13, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1038])])).
fof(f22030, plain, (~ spl42_151 | ~ spl42_494 | ~ spl42_1016 | spl42_1036), inference(avatar_contradiction_clause, [], [f22029])).
fof(f22029, plain, ($false | (~ spl42_151 | ~ spl42_494 | ~ spl42_1016 | spl42_1036)), inference(subsumption_resolution, [], [f22028, f3743])).
fof(f3743, plain, ((e23 = op2(e23, e26)) | ~ spl42_494), inference(avatar_component_clause, [], [f3741])).
fof(f3741, plain, (spl42_494 <=> (e23 = op2(e23, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_494])])).
fof(f22028, plain, (~ (e23 = op2(e23, e26)) | (~ spl42_151 | ~ spl42_1016 | spl42_1036)), inference(forward_demodulation, [], [f22027, f4800])).
fof(f22027, plain, (~ (op2(e23, e26) = h6(e13)) | (~ spl42_151 | ~ spl42_1016 | spl42_1036)), inference(forward_demodulation, [], [f22026, f2224])).
fof(f2224, plain, ((e13 = op1(e13, e16)) | ~ spl42_151), inference(avatar_component_clause, [], [f2222])).
fof(f2222, plain, (spl42_151 <=> (e13 = op1(e13, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_151])])).
fof(f22026, plain, (~ (op2(e23, e26) = h6(op1(e13, e16))) | (~ spl42_1016 | spl42_1036)), inference(forward_demodulation, [], [f6766, f6685])).
fof(f6766, plain, (~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | spl42_1036), inference(avatar_component_clause, [], [f6764])).
fof(f6764, plain, (spl42_1036 <=> (h6(op1(e13, e16)) = op2(e23, h6(e16)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1036])])).
fof(f21977, plain, (~ spl42_145 | ~ spl42_488 | spl42_1035), inference(avatar_contradiction_clause, [], [f21976])).
fof(f21976, plain, ($false | (~ spl42_145 | ~ spl42_488 | spl42_1035)), inference(subsumption_resolution, [], [f21975, f3718])).
fof(f3718, plain, ((e24 = op2(e24, e20)) | ~ spl42_488), inference(avatar_component_clause, [], [f3716])).
fof(f3716, plain, (spl42_488 <=> (e24 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_488])])).
fof(f21975, plain, (~ (e24 = op2(e24, e20)) | (~ spl42_145 | spl42_1035)), inference(forward_demodulation, [], [f21974, f4797])).
fof(f21974, plain, (~ (op2(e24, e20) = h6(e14)) | (~ spl42_145 | spl42_1035)), inference(forward_demodulation, [], [f6762, f2199])).
fof(f6762, plain, (~ (op2(e24, e20) = h6(op1(e14, e10))) | spl42_1035), inference(avatar_component_clause, [], [f6760])).
fof(f6760, plain, (spl42_1035 <=> (op2(e24, e20) = h6(op1(e14, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1035])])).
fof(f21923, plain, (~ spl42_127 | ~ spl42_785 | spl42_1033), inference(avatar_contradiction_clause, [], [f21922])).
fof(f21922, plain, ($false | (~ spl42_127 | ~ spl42_785 | spl42_1033)), inference(subsumption_resolution, [], [f21921, f5369])).
fof(f5369, plain, ((e20 = h5(e14)) | ~ spl42_785), inference(avatar_component_clause, [], [f5368])).
fof(f5368, plain, (spl42_785 <=> (e20 = h5(e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_785])])).
fof(f21921, plain, (~ (e20 = h5(e14)) | (~ spl42_127 | spl42_1033)), inference(forward_demodulation, [], [f21920, f4820])).
fof(f21920, plain, (~ (h5(e14) = h6(e10)) | (~ spl42_127 | spl42_1033)), inference(forward_demodulation, [], [f6754, f2125])).
fof(f2125, plain, ((e10 = op1(e14, e12)) | ~ spl42_127), inference(avatar_component_clause, [], [f2123])).
fof(f2123, plain, (spl42_127 <=> (e10 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_127])])).
fof(f6754, plain, (~ (h5(e14) = h6(op1(e14, e12))) | spl42_1033), inference(avatar_component_clause, [], [f6752])).
fof(f6752, plain, (spl42_1033 <=> (h5(e14) = h6(op1(e14, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1033])])).
fof(f21863, plain, (~ spl42_115 | ~ spl42_743 | spl42_1031), inference(avatar_contradiction_clause, [], [f21862])).
fof(f21862, plain, ($false | (~ spl42_115 | ~ spl42_743 | spl42_1031)), inference(subsumption_resolution, [], [f21861, f5143])).
fof(f21861, plain, (~ (e22 = h5(e11)) | (~ spl42_115 | ~ spl42_743 | spl42_1031)), inference(forward_demodulation, [], [f21860, f11296])).
fof(f21860, plain, (~ (h5(e11) = h6(e12)) | (~ spl42_115 | spl42_1031)), inference(forward_demodulation, [], [f6746, f2075])).
fof(f2075, plain, ((e12 = op1(e14, e14)) | ~ spl42_115), inference(avatar_component_clause, [], [f2073])).
fof(f2073, plain, (spl42_115 <=> (e12 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_115])])).
fof(f6746, plain, (~ (h5(e11) = h6(op1(e14, e14))) | spl42_1031), inference(avatar_component_clause, [], [f6744])).
fof(f6744, plain, (spl42_1031 <=> (h5(e11) = h6(op1(e14, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1031])])).
fof(f21800, plain, (~ spl42_89 | spl42_1027), inference(avatar_contradiction_clause, [], [f21799])).
fof(f21799, plain, ($false | (~ spl42_89 | spl42_1027)), inference(subsumption_resolution, [], [f21798, f4797])).
fof(f21798, plain, (~ (e24 = h6(e14)) | (~ spl42_89 | spl42_1027)), inference(forward_demodulation, [], [f6730, f1967])).
fof(f6730, plain, (~ (e24 = h6(op1(e15, e11))) | spl42_1027), inference(avatar_component_clause, [], [f6728])).
fof(f6728, plain, (spl42_1027 <=> (e24 = h6(op1(e15, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1027])])).
fof(f21720, plain, (~ spl42_16 | ~ spl42_359 | ~ spl42_748 | ~ spl42_1016 | spl42_1018), inference(avatar_contradiction_clause, [], [f21719])).
fof(f21719, plain, ($false | (~ spl42_16 | ~ spl42_359 | ~ spl42_748 | ~ spl42_1016 | spl42_1018)), inference(subsumption_resolution, [], [f21718, f3184])).
fof(f3184, plain, ((e21 = op2(e26, e24)) | ~ spl42_359), inference(avatar_component_clause, [], [f3182])).
fof(f3182, plain, (spl42_359 <=> (e21 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_359])])).
fof(f21718, plain, (~ (e21 = op2(e26, e24)) | (~ spl42_16 | ~ spl42_748 | ~ spl42_1016 | spl42_1018)), inference(forward_demodulation, [], [f21717, f5172])).
fof(f21717, plain, (~ (op2(e26, e24) = h6(e11)) | (~ spl42_16 | ~ spl42_1016 | spl42_1018)), inference(forward_demodulation, [], [f21716, f1665])).
fof(f1665, plain, ((e11 = op1(e16, e14)) | ~ spl42_16), inference(avatar_component_clause, [], [f1663])).
fof(f1663, plain, (spl42_16 <=> (e11 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_16])])).
fof(f21716, plain, (~ (op2(e26, e24) = h6(op1(e16, e14))) | (~ spl42_1016 | spl42_1018)), inference(forward_demodulation, [], [f6694, f6685])).
fof(f6694, plain, (~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | spl42_1018), inference(avatar_component_clause, [], [f6692])).
fof(f6692, plain, (spl42_1018 <=> (h6(op1(e16, e14)) = op2(h6(e16), e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_1018])])).
fof(f21649, plain, (~ spl42_1 | ~ spl42_344 | spl42_1015 | ~ spl42_1016), inference(avatar_contradiction_clause, [], [f21648])).
fof(f21648, plain, ($false | (~ spl42_1 | ~ spl42_344 | spl42_1015 | ~ spl42_1016)), inference(subsumption_resolution, [], [f21647, f3122])).
fof(f3122, plain, ((e20 = op2(e26, e26)) | ~ spl42_344), inference(avatar_component_clause, [], [f3120])).
fof(f3120, plain, (spl42_344 <=> (e20 = op2(e26, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_344])])).
fof(f21647, plain, (~ (e20 = op2(e26, e26)) | (~ spl42_1 | spl42_1015 | ~ spl42_1016)), inference(forward_demodulation, [], [f21646, f4820])).
fof(f21646, plain, (~ (op2(e26, e26) = h6(e10)) | (~ spl42_1 | spl42_1015 | ~ spl42_1016)), inference(forward_demodulation, [], [f21645, f1603])).
fof(f1603, plain, ((e10 = op1(e16, e16)) | ~ spl42_1), inference(avatar_component_clause, [], [f1601])).
fof(f21645, plain, (~ (op2(e26, e26) = h6(op1(e16, e16))) | (spl42_1015 | ~ spl42_1016)), inference(forward_demodulation, [], [f6682, f6685])).
fof(f6682, plain, (~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | spl42_1015), inference(avatar_component_clause, [], [f6680])).
fof(f6680, plain, (spl42_1015 <=> (h6(op1(e16, e16)) = op2(h6(e16), h6(e16)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1015])])).
fof(f21593, plain, (spl42_573 | ~ spl42_521 | ~ spl42_550), inference(avatar_split_clause, [], [f21346, f3973, f3853, f4068])).
fof(f3973, plain, (spl42_550 <=> (e23 = op2(e22, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_550])])).
fof(f21346, plain, ((e25 = op2(e22, e22)) | (~ spl42_521 | ~ spl42_550)), inference(backward_demodulation, [], [f13666, f3855])).
fof(f13666, plain, ((e25 = op2(op2(e23, e22), e22)) | ~ spl42_550), inference(forward_demodulation, [], [f491, f3975])).
fof(f3975, plain, ((e23 = op2(e22, e25)) | ~ spl42_550), inference(avatar_component_clause, [], [f3973])).
fof(f491, plain, (e25 = op2(op2(op2(e22, e25), e22), e22)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ((e26 = op2(op2(op2(e26, e26), e26), e26)) & (e26 = op2(op2(op2(e25, e26), e25), e25)) & (e26 = op2(op2(op2(e24, e26), e24), e24)) & (e26 = op2(op2(op2(e23, e26), e23), e23)) & (e26 = op2(op2(op2(e22, e26), e22), e22)) & (e26 = op2(op2(op2(e21, e26), e21), e21)) & (e26 = op2(op2(op2(e20, e26), e20), e20)) & (e25 = op2(op2(op2(e26, e25), e26), e26)) & (e25 = op2(op2(op2(e25, e25), e25), e25)) & (e25 = op2(op2(op2(e24, e25), e24), e24)) & (e25 = op2(op2(op2(e23, e25), e23), e23)) & (e25 = op2(op2(op2(e22, e25), e22), e22)) & (e25 = op2(op2(op2(e21, e25), e21), e21)) & (e25 = op2(op2(op2(e20, e25), e20), e20)) & (e24 = op2(op2(op2(e26, e24), e26), e26)) & (e24 = op2(op2(op2(e25, e24), e25), e25)) & (e24 = op2(op2(op2(e24, e24), e24), e24)) & (e24 = op2(op2(op2(e23, e24), e23), e23)) & (e24 = op2(op2(op2(e22, e24), e22), e22)) & (e24 = op2(op2(op2(e21, e24), e21), e21)) & (e24 = op2(op2(op2(e20, e24), e20), e20)) & (e23 = op2(op2(op2(e26, e23), e26), e26)) & (e23 = op2(op2(op2(e25, e23), e25), e25)) & (e23 = op2(op2(op2(e24, e23), e24), e24)) & (e23 = op2(op2(op2(e23, e23), e23), e23)) & (e23 = op2(op2(op2(e22, e23), e22), e22)) & (e23 = op2(op2(op2(e21, e23), e21), e21)) & (e23 = op2(op2(op2(e20, e23), e20), e20)) & (e22 = op2(op2(op2(e26, e22), e26), e26)) & (e22 = op2(op2(op2(e25, e22), e25), e25)) & (e22 = op2(op2(op2(e24, e22), e24), e24)) & (e22 = op2(op2(op2(e23, e22), e23), e23)) & (e22 = op2(op2(op2(e22, e22), e22), e22)) & (e22 = op2(op2(op2(e21, e22), e21), e21)) & (e22 = op2(op2(op2(e20, e22), e20), e20)) & (e21 = op2(op2(op2(e26, e21), e26), e26)) & (e21 = op2(op2(op2(e25, e21), e25), e25)) & (e21 = op2(op2(op2(e24, e21), e24), e24)) & (e21 = op2(op2(op2(e23, e21), e23), e23)) & (e21 = op2(op2(op2(e22, e21), e22), e22)) & (e21 = op2(op2(op2(e21, e21), e21), e21)) & (e21 = op2(op2(op2(e20, e21), e20), e20)) & (e20 = op2(op2(op2(e26, e20), e26), e26)) & (e20 = op2(op2(op2(e25, e20), e25), e25)) & (e20 = op2(op2(op2(e24, e20), e24), e24)) & (e20 = op2(op2(op2(e23, e20), e23), e23)) & (e20 = op2(op2(op2(e22, e20), e22), e22)) & (e20 = op2(op2(op2(e21, e20), e21), e21)) & (e20 = op2(op2(op2(e20, e20), e20), e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax6)).
fof(f21543, plain, (~ spl42_14 | ~ spl42_357 | ~ spl42_1016 | spl42_1017), inference(avatar_contradiction_clause, [], [f21542])).
fof(f21542, plain, ($false | (~ spl42_14 | ~ spl42_357 | ~ spl42_1016 | spl42_1017)), inference(subsumption_resolution, [], [f21541, f3175])).
fof(f3175, plain, ((e26 = op2(e26, e25)) | ~ spl42_357), inference(avatar_component_clause, [], [f3173])).
fof(f3173, plain, (spl42_357 <=> (e26 = op2(e26, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_357])])).
fof(f21541, plain, (~ (e26 = op2(e26, e25)) | (~ spl42_14 | ~ spl42_1016 | spl42_1017)), inference(forward_demodulation, [], [f21540, f6685])).
fof(f21540, plain, (~ (op2(e26, e25) = h6(e16)) | (~ spl42_14 | ~ spl42_1016 | spl42_1017)), inference(forward_demodulation, [], [f21539, f1656])).
fof(f1656, plain, ((e16 = op1(e16, e15)) | ~ spl42_14), inference(avatar_component_clause, [], [f1654])).
fof(f1654, plain, (spl42_14 <=> (e16 = op1(e16, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_14])])).
fof(f21539, plain, (~ (op2(e26, e25) = h6(op1(e16, e15))) | (~ spl42_1016 | spl42_1017)), inference(forward_demodulation, [], [f6690, f6685])).
fof(f6690, plain, (~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | spl42_1017), inference(avatar_component_clause, [], [f6688])).
fof(f6688, plain, (spl42_1017 <=> (h6(op1(e16, e15)) = op2(h6(e16), e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_1017])])).
fof(f21514, plain, (spl42_248 | ~ spl42_33 | ~ spl42_100), inference(avatar_split_clause, [], [f21024, f2011, f1733, f2624])).
fof(f21024, plain, ((e12 = op1(e11, e16)) | (~ spl42_33 | ~ spl42_100)), inference(forward_demodulation, [], [f21019, f2013])).
fof(f21019, plain, ((e12 = op1(op1(e14, e16), e16)) | ~ spl42_33), inference(backward_demodulation, [], [f278, f1735])).
fof(f278, plain, (e12 = op1(op1(op1(e16, e12), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f21476, plain, (~ spl42_287 | ~ spl42_630 | ~ spl42_748 | spl42_1012 | ~ spl42_1016), inference(avatar_contradiction_clause, [], [f21475])).
fof(f21475, plain, ($false | (~ spl42_287 | ~ spl42_630 | ~ spl42_748 | spl42_1012 | ~ spl42_1016)), inference(subsumption_resolution, [], [f21474, f4306])).
fof(f4306, plain, ((e26 = op2(e21, e21)) | ~ spl42_630), inference(avatar_component_clause, [], [f4304])).
fof(f4304, plain, (spl42_630 <=> (e26 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_630])])).
fof(f21474, plain, (~ (e26 = op2(e21, e21)) | (~ spl42_287 | ~ spl42_748 | spl42_1012 | ~ spl42_1016)), inference(forward_demodulation, [], [f21473, f6685])).
fof(f21473, plain, (~ (op2(e21, e21) = h6(e16)) | (~ spl42_287 | ~ spl42_748 | spl42_1012)), inference(forward_demodulation, [], [f21472, f2787])).
fof(f2787, plain, ((e16 = op1(e11, e11)) | ~ spl42_287), inference(avatar_component_clause, [], [f2785])).
fof(f2785, plain, (spl42_287 <=> (e16 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_287])])).
fof(f21472, plain, (~ (op2(e21, e21) = h6(op1(e11, e11))) | (~ spl42_748 | spl42_1012)), inference(forward_demodulation, [], [f6670, f5172])).
fof(f6670, plain, (~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | spl42_1012), inference(avatar_component_clause, [], [f6668])).
fof(f6668, plain, (spl42_1012 <=> (h6(op1(e11, e11)) = op2(h6(e11), h6(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1012])])).
fof(f21268, plain, (spl42_391 | ~ spl42_588 | ~ spl42_647), inference(avatar_split_clause, [], [f21267, f4375, f4130, f3314])).
fof(f21267, plain, ((e25 = op2(e26, e20)) | (~ spl42_588 | ~ spl42_647)), inference(forward_demodulation, [], [f21260, f4132])).
fof(f21260, plain, ((e25 = op2(op2(e22, e20), e20)) | ~ spl42_647), inference(backward_demodulation, [], [f489, f4377])).
fof(f489, plain, (e25 = op2(op2(op2(e20, e25), e20), e20)), inference(cnf_transformation, [], [f6])).
fof(f21249, plain, (~ spl42_623 | ~ spl42_672), inference(avatar_split_clause, [], [f21241, f4478, f4275])).
fof(f4275, plain, (spl42_623 <=> (e26 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_623])])).
fof(f21241, plain, (~ (e26 = op2(e21, e22)) | ~ spl42_672), inference(backward_demodulation, [], [f839, f4480])).
fof(f839, plain, ~ (op2(e20, e22) = op2(e21, e22)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (op2(e26, e25) = op2(e26, e26)) & ~ (op2(e26, e24) = op2(e26, e26)) & ~ (op2(e26, e24) = op2(e26, e25)) & ~ (op2(e26, e23) = op2(e26, e26)) & ~ (op2(e26, e23) = op2(e26, e25)) & ~ (op2(e26, e23) = op2(e26, e24)) & ~ (op2(e26, e22) = op2(e26, e26)) & ~ (op2(e26, e22) = op2(e26, e25)) & ~ (op2(e26, e22) = op2(e26, e24)) & ~ (op2(e26, e22) = op2(e26, e23)) & ~ (op2(e26, e21) = op2(e26, e26)) & ~ (op2(e26, e21) = op2(e26, e25)) & ~ (op2(e26, e21) = op2(e26, e24)) & ~ (op2(e26, e21) = op2(e26, e23)) & ~ (op2(e26, e21) = op2(e26, e22)) & ~ (op2(e26, e20) = op2(e26, e26)) & ~ (op2(e26, e20) = op2(e26, e25)) & ~ (op2(e26, e20) = op2(e26, e24)) & ~ (op2(e26, e20) = op2(e26, e23)) & ~ (op2(e26, e20) = op2(e26, e22)) & ~ (op2(e26, e20) = op2(e26, e21)) & ~ (op2(e25, e25) = op2(e25, e26)) & ~ (op2(e25, e24) = op2(e25, e26)) & ~ (op2(e25, e24) = op2(e25, e25)) & ~ (op2(e25, e23) = op2(e25, e26)) & ~ (op2(e25, e23) = op2(e25, e25)) & ~ (op2(e25, e23) = op2(e25, e24)) & ~ (op2(e25, e22) = op2(e25, e26)) & ~ (op2(e25, e22) = op2(e25, e25)) & ~ (op2(e25, e22) = op2(e25, e24)) & ~ (op2(e25, e22) = op2(e25, e23)) & ~ (op2(e25, e21) = op2(e25, e26)) & ~ (op2(e25, e21) = op2(e25, e25)) & ~ (op2(e25, e21) = op2(e25, e24)) & ~ (op2(e25, e21) = op2(e25, e23)) & ~ (op2(e25, e21) = op2(e25, e22)) & ~ (op2(e25, e20) = op2(e25, e26)) & ~ (op2(e25, e20) = op2(e25, e25)) & ~ (op2(e25, e20) = op2(e25, e24)) & ~ (op2(e25, e20) = op2(e25, e23)) & ~ (op2(e25, e20) = op2(e25, e22)) & ~ (op2(e25, e20) = op2(e25, e21)) & ~ (op2(e24, e25) = op2(e24, e26)) & ~ (op2(e24, e24) = op2(e24, e26)) & ~ (op2(e24, e24) = op2(e24, e25)) & ~ (op2(e24, e23) = op2(e24, e26)) & ~ (op2(e24, e23) = op2(e24, e25)) & ~ (op2(e24, e23) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e26)) & ~ (op2(e24, e22) = op2(e24, e25)) & ~ (op2(e24, e22) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e26)) & ~ (op2(e24, e21) = op2(e24, e25)) & ~ (op2(e24, e21) = op2(e24, e24)) & ~ (op2(e24, e21) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e26)) & ~ (op2(e24, e20) = op2(e24, e25)) & ~ (op2(e24, e20) = op2(e24, e24)) & ~ (op2(e24, e20) = op2(e24, e23)) & ~ (op2(e24, e20) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e21)) & ~ (op2(e23, e25) = op2(e23, e26)) & ~ (op2(e23, e24) = op2(e23, e26)) & ~ (op2(e23, e24) = op2(e23, e25)) & ~ (op2(e23, e23) = op2(e23, e26)) & ~ (op2(e23, e23) = op2(e23, e25)) & ~ (op2(e23, e23) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e26)) & ~ (op2(e23, e22) = op2(e23, e25)) & ~ (op2(e23, e22) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e26)) & ~ (op2(e23, e21) = op2(e23, e25)) & ~ (op2(e23, e21) = op2(e23, e24)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e26)) & ~ (op2(e23, e20) = op2(e23, e25)) & ~ (op2(e23, e20) = op2(e23, e24)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e25) = op2(e22, e26)) & ~ (op2(e22, e24) = op2(e22, e26)) & ~ (op2(e22, e24) = op2(e22, e25)) & ~ (op2(e22, e23) = op2(e22, e26)) & ~ (op2(e22, e23) = op2(e22, e25)) & ~ (op2(e22, e23) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e26)) & ~ (op2(e22, e22) = op2(e22, e25)) & ~ (op2(e22, e22) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e26)) & ~ (op2(e22, e21) = op2(e22, e25)) & ~ (op2(e22, e21) = op2(e22, e24)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e26)) & ~ (op2(e22, e20) = op2(e22, e25)) & ~ (op2(e22, e20) = op2(e22, e24)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e25) = op2(e21, e26)) & ~ (op2(e21, e24) = op2(e21, e26)) & ~ (op2(e21, e24) = op2(e21, e25)) & ~ (op2(e21, e23) = op2(e21, e26)) & ~ (op2(e21, e23) = op2(e21, e25)) & ~ (op2(e21, e23) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e26)) & ~ (op2(e21, e22) = op2(e21, e25)) & ~ (op2(e21, e22) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e26)) & ~ (op2(e21, e21) = op2(e21, e25)) & ~ (op2(e21, e21) = op2(e21, e24)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e26)) & ~ (op2(e21, e20) = op2(e21, e25)) & ~ (op2(e21, e20) = op2(e21, e24)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e25) = op2(e20, e26)) & ~ (op2(e20, e24) = op2(e20, e26)) & ~ (op2(e20, e24) = op2(e20, e25)) & ~ (op2(e20, e23) = op2(e20, e26)) & ~ (op2(e20, e23) = op2(e20, e25)) & ~ (op2(e20, e23) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e26)) & ~ (op2(e20, e22) = op2(e20, e25)) & ~ (op2(e20, e22) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e26)) & ~ (op2(e20, e21) = op2(e20, e25)) & ~ (op2(e20, e21) = op2(e20, e24)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e26)) & ~ (op2(e20, e20) = op2(e20, e25)) & ~ (op2(e20, e20) = op2(e20, e24)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e25, e26) = op2(e26, e26)) & ~ (op2(e24, e26) = op2(e26, e26)) & ~ (op2(e24, e26) = op2(e25, e26)) & ~ (op2(e23, e26) = op2(e26, e26)) & ~ (op2(e23, e26) = op2(e25, e26)) & ~ (op2(e23, e26) = op2(e24, e26)) & ~ (op2(e22, e26) = op2(e26, e26)) & ~ (op2(e22, e26) = op2(e25, e26)) & ~ (op2(e22, e26) = op2(e24, e26)) & ~ (op2(e22, e26) = op2(e23, e26)) & ~ (op2(e21, e26) = op2(e26, e26)) & ~ (op2(e21, e26) = op2(e25, e26)) & ~ (op2(e21, e26) = op2(e24, e26)) & ~ (op2(e21, e26) = op2(e23, e26)) & ~ (op2(e21, e26) = op2(e22, e26)) & ~ (op2(e20, e26) = op2(e26, e26)) & ~ (op2(e20, e26) = op2(e25, e26)) & ~ (op2(e20, e26) = op2(e24, e26)) & ~ (op2(e20, e26) = op2(e23, e26)) & ~ (op2(e20, e26) = op2(e22, e26)) & ~ (op2(e20, e26) = op2(e21, e26)) & ~ (op2(e25, e25) = op2(e26, e25)) & ~ (op2(e24, e25) = op2(e26, e25)) & ~ (op2(e24, e25) = op2(e25, e25)) & ~ (op2(e23, e25) = op2(e26, e25)) & ~ (op2(e23, e25) = op2(e25, e25)) & ~ (op2(e23, e25) = op2(e24, e25)) & ~ (op2(e22, e25) = op2(e26, e25)) & ~ (op2(e22, e25) = op2(e25, e25)) & ~ (op2(e22, e25) = op2(e24, e25)) & ~ (op2(e22, e25) = op2(e23, e25)) & ~ (op2(e21, e25) = op2(e26, e25)) & ~ (op2(e21, e25) = op2(e25, e25)) & ~ (op2(e21, e25) = op2(e24, e25)) & ~ (op2(e21, e25) = op2(e23, e25)) & ~ (op2(e21, e25) = op2(e22, e25)) & ~ (op2(e20, e25) = op2(e26, e25)) & ~ (op2(e20, e25) = op2(e25, e25)) & ~ (op2(e20, e25) = op2(e24, e25)) & ~ (op2(e20, e25) = op2(e23, e25)) & ~ (op2(e20, e25) = op2(e22, e25)) & ~ (op2(e20, e25) = op2(e21, e25)) & ~ (op2(e25, e24) = op2(e26, e24)) & ~ (op2(e24, e24) = op2(e26, e24)) & ~ (op2(e24, e24) = op2(e25, e24)) & ~ (op2(e23, e24) = op2(e26, e24)) & ~ (op2(e23, e24) = op2(e25, e24)) & ~ (op2(e23, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e26, e24)) & ~ (op2(e22, e24) = op2(e25, e24)) & ~ (op2(e22, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e26, e24)) & ~ (op2(e21, e24) = op2(e25, e24)) & ~ (op2(e21, e24) = op2(e24, e24)) & ~ (op2(e21, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e26, e24)) & ~ (op2(e20, e24) = op2(e25, e24)) & ~ (op2(e20, e24) = op2(e24, e24)) & ~ (op2(e20, e24) = op2(e23, e24)) & ~ (op2(e20, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e21, e24)) & ~ (op2(e25, e23) = op2(e26, e23)) & ~ (op2(e24, e23) = op2(e26, e23)) & ~ (op2(e24, e23) = op2(e25, e23)) & ~ (op2(e23, e23) = op2(e26, e23)) & ~ (op2(e23, e23) = op2(e25, e23)) & ~ (op2(e23, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e26, e23)) & ~ (op2(e22, e23) = op2(e25, e23)) & ~ (op2(e22, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e26, e23)) & ~ (op2(e21, e23) = op2(e25, e23)) & ~ (op2(e21, e23) = op2(e24, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e26, e23)) & ~ (op2(e20, e23) = op2(e25, e23)) & ~ (op2(e20, e23) = op2(e24, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e25, e22) = op2(e26, e22)) & ~ (op2(e24, e22) = op2(e26, e22)) & ~ (op2(e24, e22) = op2(e25, e22)) & ~ (op2(e23, e22) = op2(e26, e22)) & ~ (op2(e23, e22) = op2(e25, e22)) & ~ (op2(e23, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e26, e22)) & ~ (op2(e22, e22) = op2(e25, e22)) & ~ (op2(e22, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e26, e22)) & ~ (op2(e21, e22) = op2(e25, e22)) & ~ (op2(e21, e22) = op2(e24, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e26, e22)) & ~ (op2(e20, e22) = op2(e25, e22)) & ~ (op2(e20, e22) = op2(e24, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e25, e21) = op2(e26, e21)) & ~ (op2(e24, e21) = op2(e26, e21)) & ~ (op2(e24, e21) = op2(e25, e21)) & ~ (op2(e23, e21) = op2(e26, e21)) & ~ (op2(e23, e21) = op2(e25, e21)) & ~ (op2(e23, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e26, e21)) & ~ (op2(e22, e21) = op2(e25, e21)) & ~ (op2(e22, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e26, e21)) & ~ (op2(e21, e21) = op2(e25, e21)) & ~ (op2(e21, e21) = op2(e24, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e26, e21)) & ~ (op2(e20, e21) = op2(e25, e21)) & ~ (op2(e20, e21) = op2(e24, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e25, e20) = op2(e26, e20)) & ~ (op2(e24, e20) = op2(e26, e20)) & ~ (op2(e24, e20) = op2(e25, e20)) & ~ (op2(e23, e20) = op2(e26, e20)) & ~ (op2(e23, e20) = op2(e25, e20)) & ~ (op2(e23, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e26, e20)) & ~ (op2(e22, e20) = op2(e25, e20)) & ~ (op2(e22, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e26, e20)) & ~ (op2(e21, e20) = op2(e25, e20)) & ~ (op2(e21, e20) = op2(e24, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e26, e20)) & ~ (op2(e20, e20) = op2(e25, e20)) & ~ (op2(e20, e20) = op2(e24, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax8)).
fof(f21238, plain, (spl42_534 | ~ spl42_673 | ~ spl42_683), inference(avatar_split_clause, [], [f21237, f4524, f4483, f3907])).
fof(f4524, plain, (spl42_683 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl42_683])])).
fof(f21237, plain, ((e21 = op2(e23, e20)) | (~ spl42_673 | ~ spl42_683)), inference(forward_demodulation, [], [f21227, f4526])).
fof(f4526, plain, ((op2(e20, e20) = e23) | ~ spl42_683), inference(avatar_component_clause, [], [f4524])).
fof(f21227, plain, ((e21 = op2(op2(e20, e20), e20)) | ~ spl42_673), inference(backward_demodulation, [], [f461, f4485])).
fof(f461, plain, (e21 = op2(op2(op2(e20, e21), e20), e20)), inference(cnf_transformation, [], [f6])).
fof(f21093, plain, (~ spl42_414 | ~ spl42_71 | spl42_1025), inference(avatar_split_clause, [], [f21092, f6720, f1891, f3410])).
fof(f3410, plain, (spl42_414 <=> (e20 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_414])])).
fof(f1891, plain, (spl42_71 <=> (e10 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_71])])).
fof(f6720, plain, (spl42_1025 <=> (op2(e25, e23) = h6(op1(e15, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1025])])).
fof(f21092, plain, (~ (e20 = op2(e25, e23)) | (~ spl42_71 | spl42_1025)), inference(forward_demodulation, [], [f21091, f4820])).
fof(f21091, plain, (~ (op2(e25, e23) = h6(e10)) | (~ spl42_71 | spl42_1025)), inference(forward_demodulation, [], [f6722, f1893])).
fof(f1893, plain, ((e10 = op1(e15, e13)) | ~ spl42_71), inference(avatar_component_clause, [], [f1891])).
fof(f6722, plain, (~ (op2(e25, e23) = h6(op1(e15, e13))) | spl42_1025), inference(avatar_component_clause, [], [f6720])).
fof(f21042, plain, (~ spl42_389 | ~ spl42_368), inference(avatar_split_clause, [], [f21041, f3219, f3306])).
fof(f3306, plain, (spl42_389 <=> (e23 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_389])])).
fof(f21041, plain, (~ (e23 = op2(e26, e20)) | ~ spl42_368), inference(forward_demodulation, [], [f1072, f3221])).
fof(f1072, plain, ~ (op2(e26, e20) = op2(e26, e23)), inference(cnf_transformation, [], [f8])).
fof(f21001, plain, (spl42_1 | ~ spl42_48 | ~ spl42_56), inference(avatar_split_clause, [], [f21000, f1828, f1795, f1601])).
fof(f21000, plain, ((e10 = op1(e16, e16)) | (~ spl42_48 | ~ spl42_56)), inference(forward_demodulation, [], [f20994, f1830])).
fof(f20994, plain, ((e10 = op1(op1(e15, e16), e16)) | ~ spl42_48), inference(backward_demodulation, [], [f264, f1797])).
fof(f264, plain, (e10 = op1(op1(op1(e16, e10), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f20993, plain, (~ spl42_170 | ~ spl42_173), inference(avatar_contradiction_clause, [], [f20992])).
fof(f20992, plain, ($false | (~ spl42_170 | ~ spl42_173)), inference(subsumption_resolution, [], [f20991, f1099])).
fof(f1099, plain, ~ (e11 = e14), inference(cnf_transformation, [], [f9])).
fof(f20991, plain, ((e11 = e14) | (~ spl42_170 | ~ spl42_173)), inference(backward_demodulation, [], [f2315, f2303])).
fof(f2303, plain, ((e11 = op1(e13, e13)) | ~ spl42_170), inference(avatar_component_clause, [], [f2301])).
fof(f2301, plain, (spl42_170 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_170])])).
fof(f20988, plain, (~ spl42_31 | ~ spl42_178), inference(avatar_split_clause, [], [f20981, f2334, f1725])).
fof(f20981, plain, (~ (e12 = op1(e16, e12)) | ~ spl42_178), inference(backward_demodulation, [], [f562, f2336])).
fof(f562, plain, ~ (op1(e13, e12) = op1(e16, e12)), inference(cnf_transformation, [], [f7])).
fof(f20973, plain, (~ spl42_202 | ~ spl42_230), inference(avatar_split_clause, [], [f20969, f2549, f2433])).
fof(f2433, plain, (spl42_202 <=> (e15 = op1(e12, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_202])])).
fof(f20969, plain, (~ (e15 = op1(e12, e16)) | ~ spl42_230), inference(backward_demodulation, [], [f706, f2551])).
fof(f706, plain, ~ (op1(e12, e12) = op1(e12, e16)), inference(cnf_transformation, [], [f7])).
fof(f20906, plain, (spl42_800 | ~ spl42_348 | ~ spl42_1168), inference(avatar_split_clause, [], [f20905, f7615, f3136, f5444])).
fof(f5444, plain, (spl42_800 <=> (e24 = h4(e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_800])])).
fof(f3136, plain, (spl42_348 <=> (e24 = op2(e26, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_348])])).
fof(f7615, plain, (spl42_1168 <=> (e26 = h4(e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_1168])])).
fof(f20905, plain, ((e24 = h4(e12)) | (~ spl42_348 | ~ spl42_1168)), inference(backward_demodulation, [], [f19985, f3138])).
fof(f3138, plain, ((e24 = op2(e26, e26)) | ~ spl42_348), inference(avatar_component_clause, [], [f3136])).
fof(f19985, plain, ((op2(e26, e26) = h4(e12)) | ~ spl42_1168), inference(backward_demodulation, [], [f4741, f7616])).
fof(f7616, plain, ((e26 = h4(e14)) | ~ spl42_1168), inference(avatar_component_clause, [], [f7615])).
fof(f4741, plain, (h4(e12) = op2(h4(e14), h4(e14))), inference(forward_demodulation, [], [f1232, f1234])).
fof(f1234, plain, (op2(e23, op2(e23, e23)) = h4(e14)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ((h4(e16) = op2(op2(op2(op2(e23, op2(e23, e23)), op2(e23, op2(e23, e23))), e23), op2(e23, op2(e23, e23)))) & (op2(e23, op2(e23, e23)) = h4(e14)) & (h4(e13) = op2(op2(op2(e23, op2(e23, e23)), op2(e23, op2(e23, e23))), e23)) & (op2(op2(e23, op2(e23, e23)), op2(e23, op2(e23, e23))) = h4(e12)) & (op2(e23, e23) = h4(e11)) & (h4(e10) = op2(op2(op2(e23, op2(e23, e23)), op2(e23, op2(e23, e23))), op2(e23, op2(e23, e23)))) & (e23 = h4(e15))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax19)).
fof(f1232, plain, (op2(op2(e23, op2(e23, e23)), op2(e23, op2(e23, e23))) = h4(e12)), inference(cnf_transformation, [], [f19])).
fof(f20828, plain, (~ spl42_428 | ~ spl42_432), inference(avatar_contradiction_clause, [], [f20827])).
fof(f20827, plain, ($false | (~ spl42_428 | ~ spl42_432)), inference(subsumption_resolution, [], [f20824, f1115])).
fof(f1115, plain, ~ (e20 = e24), inference(cnf_transformation, [], [f10])).
fof(f10, plain, (~ (e25 = e26) & ~ (e24 = e26) & ~ (e24 = e25) & ~ (e23 = e26) & ~ (e23 = e25) & ~ (e23 = e24) & ~ (e22 = e26) & ~ (e22 = e25) & ~ (e22 = e24) & ~ (e22 = e23) & ~ (e21 = e26) & ~ (e21 = e25) & ~ (e21 = e24) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e26) & ~ (e20 = e25) & ~ (e20 = e24) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax10)).
fof(f20824, plain, ((e20 = e24) | (~ spl42_428 | ~ spl42_432)), inference(backward_demodulation, [], [f3486, f3470])).
fof(f3470, plain, ((e20 = op2(e25, e21)) | ~ spl42_428), inference(avatar_component_clause, [], [f3468])).
fof(f3468, plain, (spl42_428 <=> (e20 = op2(e25, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_428])])).
fof(f3486, plain, ((e24 = op2(e25, e21)) | ~ spl42_432), inference(avatar_component_clause, [], [f3484])).
fof(f3484, plain, (spl42_432 <=> (e24 = op2(e25, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_432])])).
fof(f20727, plain, (~ spl42_399 | ~ spl42_595), inference(avatar_split_clause, [], [f20720, f4159, f3347])).
fof(f4159, plain, (spl42_595 <=> (e26 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_595])])).
fof(f20720, plain, (~ (e26 = op2(e25, e26)) | ~ spl42_595), inference(backward_demodulation, [], [f932, f4161])).
fof(f4161, plain, ((e26 = op2(e21, e26)) | ~ spl42_595), inference(avatar_component_clause, [], [f4159])).
fof(f932, plain, ~ (op2(e21, e26) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f20657, plain, (~ spl42_647 | ~ spl42_668), inference(avatar_split_clause, [], [f20642, f4462, f4375])).
fof(f4462, plain, (spl42_668 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_668])])).
fof(f20642, plain, (~ (e22 = op2(e20, e25)) | ~ spl42_668), inference(backward_demodulation, [], [f957, f4464])).
fof(f4464, plain, ((e22 = op2(e20, e22)) | ~ spl42_668), inference(avatar_component_clause, [], [f4462])).
fof(f957, plain, ~ (op2(e20, e22) = op2(e20, e25)), inference(cnf_transformation, [], [f8])).
fof(f20612, plain, (spl42_668 | ~ spl42_470 | ~ spl42_843), inference(avatar_split_clause, [], [f20611, f5660, f3642, f4462])).
fof(f3642, plain, (spl42_470 <=> (e20 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_470])])).
fof(f5660, plain, (spl42_843 <=> (e24 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_843])])).
fof(f20611, plain, ((e22 = op2(e20, e22)) | (~ spl42_470 | ~ spl42_843)), inference(forward_demodulation, [], [f20601, f3644])).
fof(f3644, plain, ((e20 = op2(e24, e22)) | ~ spl42_470), inference(avatar_component_clause, [], [f3642])).
fof(f20601, plain, ((e22 = op2(op2(e24, e22), e22)) | ~ spl42_843), inference(backward_demodulation, [], [f4721, f5661])).
fof(f5661, plain, ((e24 = h3(e11)) | ~ spl42_843), inference(avatar_component_clause, [], [f5660])).
fof(f4721, plain, (e22 = op2(op2(h3(e11), e22), e22)), inference(backward_demodulation, [], [f470, f1224])).
fof(f1224, plain, (op2(e22, e22) = h3(e11)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ((h3(e16) = op2(op2(op2(op2(e22, op2(e22, e22)), op2(e22, op2(e22, e22))), e22), op2(e22, op2(e22, e22)))) & (op2(e22, op2(e22, e22)) = h3(e14)) & (h3(e13) = op2(op2(op2(e22, op2(e22, e22)), op2(e22, op2(e22, e22))), e22)) & (op2(op2(e22, op2(e22, e22)), op2(e22, op2(e22, e22))) = h3(e12)) & (op2(e22, e22) = h3(e11)) & (h3(e10) = op2(op2(op2(e22, op2(e22, e22)), op2(e22, op2(e22, e22))), op2(e22, op2(e22, e22)))) & (e22 = h3(e15))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax18)).
fof(f470, plain, (e22 = op2(op2(op2(e22, e22), e22), e22)), inference(cnf_transformation, [], [f6])).
fof(f20585, plain, (~ spl42_591 | ~ spl42_899), inference(avatar_split_clause, [], [f20572, f5941, f4143])).
fof(f5941, plain, (spl42_899 <=> (e22 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_899])])).
fof(f20572, plain, (~ (e22 = op2(e21, e26)) | ~ spl42_899), inference(backward_demodulation, [], [f4711, f5942])).
fof(f5942, plain, ((e22 = h2(e11)) | ~ spl42_899), inference(avatar_component_clause, [], [f5941])).
fof(f4711, plain, ~ (op2(e21, e26) = h2(e11)), inference(backward_demodulation, [], [f975, f1217])).
fof(f1217, plain, (op2(e21, e21) = h2(e11)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((h2(e16) = op2(op2(op2(op2(e21, op2(e21, e21)), op2(e21, op2(e21, e21))), e21), op2(e21, op2(e21, e21)))) & (op2(e21, op2(e21, e21)) = h2(e14)) & (h2(e13) = op2(op2(op2(e21, op2(e21, e21)), op2(e21, op2(e21, e21))), e21)) & (op2(op2(e21, op2(e21, e21)), op2(e21, op2(e21, e21))) = h2(e12)) & (op2(e21, e21) = h2(e11)) & (h2(e10) = op2(op2(op2(e21, op2(e21, e21)), op2(e21, op2(e21, e21))), op2(e21, op2(e21, e21)))) & (e21 = h2(e15))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax17)).
fof(f975, plain, ~ (op2(e21, e21) = op2(e21, e26)), inference(cnf_transformation, [], [f8])).
fof(f20505, plain, (~ spl42_635 | ~ spl42_600), inference(avatar_split_clause, [], [f18383, f4180, f4325])).
fof(f4325, plain, (spl42_635 <=> (e24 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_635])])).
fof(f18383, plain, (~ (e24 = op2(e21, e20)) | ~ spl42_600), inference(forward_demodulation, [], [f969, f4182])).
fof(f969, plain, ~ (op2(e21, e20) = op2(e21, e25)), inference(cnf_transformation, [], [f8])).
fof(f20492, plain, (~ spl42_562 | ~ spl42_576), inference(avatar_split_clause, [], [f20491, f4081, f4023])).
fof(f4023, plain, (spl42_562 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_562])])).
fof(f20491, plain, (~ (e21 = op2(e22, e23)) | ~ spl42_576), inference(forward_demodulation, [], [f993, f4083])).
fof(f993, plain, ~ (op2(e22, e21) = op2(e22, e23)), inference(cnf_transformation, [], [f8])).
fof(f20472, plain, (spl42_469 | ~ spl42_494 | ~ spl42_516), inference(avatar_split_clause, [], [f20166, f3832, f3741, f3637])).
fof(f3832, plain, (spl42_516 <=> (e24 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_516])])).
fof(f20166, plain, ((e26 = op2(e24, e23)) | (~ spl42_494 | ~ spl42_516)), inference(forward_demodulation, [], [f20159, f3834])).
fof(f3834, plain, ((e24 = op2(e23, e23)) | ~ spl42_516), inference(avatar_component_clause, [], [f3832])).
fof(f20159, plain, ((e26 = op2(op2(e23, e23), e23)) | ~ spl42_494), inference(backward_demodulation, [], [f499, f3743])).
fof(f499, plain, (e26 = op2(op2(op2(e23, e26), e23), e23)), inference(cnf_transformation, [], [f6])).
fof(f20471, plain, (~ spl42_447 | ~ spl42_454), inference(avatar_split_clause, [], [f19266, f3575, f3546])).
fof(f3546, plain, (spl42_447 <=> (e25 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_447])])).
fof(f19266, plain, (~ (e25 = op2(e24, e26)) | ~ spl42_454), inference(forward_demodulation, [], [f1048, f3577])).
fof(f1048, plain, ~ (op2(e24, e25) = op2(e24, e26)), inference(cnf_transformation, [], [f8])).
fof(f20470, plain, (~ spl42_445 | ~ spl42_494), inference(avatar_split_clause, [], [f20469, f3741, f3538])).
fof(f3538, plain, (spl42_445 <=> (e23 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_445])])).
fof(f20469, plain, (~ (e23 = op2(e24, e26)) | ~ spl42_494), inference(forward_demodulation, [], [f938, f3743])).
fof(f938, plain, ~ (op2(e23, e26) = op2(e24, e26)), inference(cnf_transformation, [], [f8])).
fof(f20462, plain, (~ spl42_396 | ~ spl42_494), inference(avatar_split_clause, [], [f20160, f3741, f3335])).
fof(f3335, plain, (spl42_396 <=> (e23 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_396])])).
fof(f20160, plain, (~ (e23 = op2(e25, e26)) | ~ spl42_494), inference(backward_demodulation, [], [f939, f3743])).
fof(f939, plain, ~ (op2(e23, e26) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f20460, plain, (~ spl42_375 | ~ spl42_368), inference(avatar_split_clause, [], [f20459, f3219, f3248])).
fof(f3248, plain, (spl42_375 <=> (e23 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_375])])).
fof(f20459, plain, (~ (e23 = op2(e26, e22)) | ~ spl42_368), inference(forward_demodulation, [], [f1081, f3221])).
fof(f1081, plain, ~ (op2(e26, e22) = op2(e26, e23)), inference(cnf_transformation, [], [f8])).
fof(f20456, plain, (spl42_48 | ~ spl42_245 | ~ spl42_304), inference(avatar_split_clause, [], [f20455, f2856, f2611, f1795])).
fof(f20455, plain, ((e15 = op1(e16, e10)) | (~ spl42_245 | ~ spl42_304)), inference(forward_demodulation, [], [f19726, f2613])).
fof(f19726, plain, ((e15 = op1(op1(e12, e10), e10)) | ~ spl42_304), inference(backward_demodulation, [], [f293, f2858])).
fof(f293, plain, (e15 = op1(op1(op1(e10, e15), e10), e10)), inference(cnf_transformation, [], [f3])).
fof(f20377, plain, (~ spl42_177 | ~ spl42_207 | spl42_279), inference(avatar_contradiction_clause, [], [f20376])).
fof(f20376, plain, ($false | (~ spl42_177 | ~ spl42_207 | spl42_279)), inference(subsumption_resolution, [], [f20371, f2753])).
fof(f2753, plain, (~ (e15 = op1(e11, e12)) | spl42_279), inference(avatar_component_clause, [], [f2752])).
fof(f2752, plain, (spl42_279 <=> (e15 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_279])])).
fof(f20371, plain, ((e15 = op1(e11, e12)) | (~ spl42_177 | ~ spl42_207)), inference(backward_demodulation, [], [f13407, f2332])).
fof(f2332, plain, ((e11 = op1(e13, e12)) | ~ spl42_177), inference(avatar_component_clause, [], [f2330])).
fof(f2330, plain, (spl42_177 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_177])])).
fof(f13407, plain, ((e15 = op1(op1(e13, e12), e12)) | ~ spl42_207), inference(forward_demodulation, [], [f295, f2456])).
fof(f295, plain, (e15 = op1(op1(op1(e12, e15), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f20317, plain, (~ spl42_233 | spl42_282 | ~ spl42_283), inference(avatar_contradiction_clause, [], [f20316])).
fof(f20316, plain, ($false | (~ spl42_233 | spl42_282 | ~ spl42_283)), inference(subsumption_resolution, [], [f20315, f2766])).
fof(f2766, plain, (~ (e11 = op1(e11, e11)) | spl42_282), inference(avatar_component_clause, [], [f2765])).
fof(f2765, plain, (spl42_282 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_282])])).
fof(f20315, plain, ((e11 = op1(e11, e11)) | (~ spl42_233 | ~ spl42_283)), inference(forward_demodulation, [], [f20309, f2564])).
fof(f20309, plain, ((e11 = op1(op1(e12, e11), e11)) | ~ spl42_283), inference(backward_demodulation, [], [f266, f2771])).
fof(f2771, plain, ((e12 = op1(e11, e11)) | ~ spl42_283), inference(avatar_component_clause, [], [f2769])).
fof(f2769, plain, (spl42_283 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_283])])).
fof(f266, plain, (e11 = op1(op1(op1(e11, e11), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f20280, plain, (~ spl42_178 | ~ spl42_325), inference(avatar_split_clause, [], [f20269, f2943, f2334])).
fof(f2943, plain, (spl42_325 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_325])])).
fof(f20269, plain, (~ (e12 = op1(e13, e12)) | ~ spl42_325), inference(backward_demodulation, [], [f547, f2945])).
fof(f2945, plain, ((e12 = op1(e10, e12)) | ~ spl42_325), inference(avatar_component_clause, [], [f2943])).
fof(f547, plain, ~ (op1(e10, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f20212, plain, (~ spl42_800 | ~ spl42_376 | ~ spl42_1168), inference(avatar_split_clause, [], [f20209, f7615, f3252, f5444])).
fof(f20209, plain, (~ (e24 = h4(e12)) | (~ spl42_376 | ~ spl42_1168)), inference(backward_demodulation, [], [f19993, f3254])).
fof(f19993, plain, (~ (op2(e26, e22) = h4(e12)) | ~ spl42_1168), inference(backward_demodulation, [], [f4835, f19987])).
fof(f19987, plain, ((h4(e12) = h7(e11)) | ~ spl42_1168), inference(backward_demodulation, [], [f1252, f19985])).
fof(f1252, plain, (op2(e26, e26) = h7(e11)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ((h7(e16) = op2(op2(op2(op2(e26, op2(e26, e26)), op2(e26, op2(e26, e26))), e26), op2(e26, op2(e26, e26)))) & (op2(e26, op2(e26, e26)) = h7(e14)) & (h7(e13) = op2(op2(op2(e26, op2(e26, e26)), op2(e26, op2(e26, e26))), e26)) & (op2(op2(e26, op2(e26, e26)), op2(e26, op2(e26, e26))) = h7(e12)) & (op2(e26, e26) = h7(e11)) & (h7(e10) = op2(op2(op2(e26, op2(e26, e26)), op2(e26, op2(e26, e26))), op2(e26, op2(e26, e26)))) & (e26 = h7(e15))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax22)).
fof(f4835, plain, ~ (op2(e26, e22) = h7(e11)), inference(backward_demodulation, [], [f1084, f1252])).
fof(f1084, plain, ~ (op2(e26, e22) = op2(e26, e26)), inference(cnf_transformation, [], [f8])).
fof(f20096, plain, (~ spl42_386 | spl42_393 | ~ spl42_643), inference(avatar_contradiction_clause, [], [f20095])).
fof(f20095, plain, ($false | (~ spl42_386 | spl42_393 | ~ spl42_643)), inference(subsumption_resolution, [], [f20090, f3324])).
fof(f3324, plain, (~ (e20 = op2(e25, e26)) | spl42_393), inference(avatar_component_clause, [], [f3323])).
fof(f3323, plain, (spl42_393 <=> (e20 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_393])])).
fof(f20090, plain, ((e20 = op2(e25, e26)) | (~ spl42_386 | ~ spl42_643)), inference(backward_demodulation, [], [f19884, f4360])).
fof(f19884, plain, ((e20 = op2(op2(e20, e26), e26)) | ~ spl42_386), inference(forward_demodulation, [], [f460, f3296])).
fof(f3296, plain, ((e20 = op2(e26, e20)) | ~ spl42_386), inference(avatar_component_clause, [], [f3294])).
fof(f3294, plain, (spl42_386 <=> (e20 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_386])])).
fof(f460, plain, (e20 = op2(op2(op2(e26, e20), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f20083, plain, (~ spl42_647 | ~ spl42_675), inference(avatar_contradiction_clause, [], [f20082])).
fof(f20082, plain, ($false | (~ spl42_647 | ~ spl42_675)), inference(subsumption_resolution, [], [f20081, f1121])).
fof(f1121, plain, ~ (e21 = e25), inference(cnf_transformation, [], [f10])).
fof(f20081, plain, ((e21 = e25) | (~ spl42_647 | ~ spl42_675)), inference(forward_demodulation, [], [f20072, f19952])).
fof(f19952, plain, ((e21 = op2(op2(e22, e20), e20)) | ~ spl42_675), inference(forward_demodulation, [], [f461, f4493])).
fof(f4493, plain, ((e22 = op2(e20, e21)) | ~ spl42_675), inference(avatar_component_clause, [], [f4491])).
fof(f4491, plain, (spl42_675 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_675])])).
fof(f20072, plain, ((e25 = op2(op2(e22, e20), e20)) | ~ spl42_647), inference(backward_demodulation, [], [f489, f4377])).
fof(f20071, plain, (~ spl42_652 | ~ spl42_656), inference(avatar_contradiction_clause, [], [f20070])).
fof(f20070, plain, ($false | (~ spl42_652 | ~ spl42_656)), inference(subsumption_resolution, [], [f20069, f1115])).
fof(f20069, plain, ((e20 = e24) | (~ spl42_652 | ~ spl42_656)), inference(backward_demodulation, [], [f4414, f4398])).
fof(f4398, plain, ((e20 = op2(e20, e24)) | ~ spl42_652), inference(avatar_component_clause, [], [f4396])).
fof(f4396, plain, (spl42_652 <=> (e20 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_652])])).
fof(f19956, plain, (~ spl42_371 | ~ spl42_469), inference(avatar_split_clause, [], [f19955, f3637, f3231])).
fof(f3231, plain, (spl42_371 <=> (e26 = op2(e26, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_371])])).
fof(f19955, plain, (~ (e26 = op2(e26, e23)) | ~ spl42_469), inference(forward_demodulation, [], [f879, f3639])).
fof(f879, plain, ~ (op2(e24, e23) = op2(e26, e23)), inference(cnf_transformation, [], [f8])).
fof(f19932, plain, (spl42_672 | ~ spl42_470 | ~ spl42_544), inference(avatar_split_clause, [], [f19516, f3948, f3642, f4478])).
fof(f19516, plain, ((e26 = op2(e20, e22)) | (~ spl42_470 | ~ spl42_544)), inference(forward_demodulation, [], [f19510, f3644])).
fof(f19510, plain, ((e26 = op2(op2(e24, e22), e22)) | ~ spl42_544), inference(backward_demodulation, [], [f498, f3950])).
fof(f498, plain, (e26 = op2(op2(op2(e22, e26), e22), e22)), inference(cnf_transformation, [], [f6])).
fof(f19929, plain, (spl42_660 | ~ spl42_414 | ~ spl42_531), inference(avatar_split_clause, [], [f19617, f3894, f3410, f4429])).
fof(f19617, plain, ((e21 = op2(e20, e23)) | (~ spl42_414 | ~ spl42_531)), inference(backward_demodulation, [], [f19376, f3412])).
fof(f3412, plain, ((e20 = op2(e25, e23)) | ~ spl42_414), inference(avatar_component_clause, [], [f3410])).
fof(f19376, plain, ((e21 = op2(op2(e25, e23), e23)) | ~ spl42_531), inference(forward_demodulation, [], [f464, f3896])).
fof(f464, plain, (e21 = op2(op2(op2(e23, e21), e23), e23)), inference(cnf_transformation, [], [f6])).
fof(f19928, plain, (spl42_647 | ~ spl42_424 | ~ spl42_498), inference(avatar_split_clause, [], [f19598, f3758, f3451, f4375])).
fof(f19598, plain, ((e22 = op2(e20, e25)) | (~ spl42_424 | ~ spl42_498)), inference(forward_demodulation, [], [f19590, f3760])).
fof(f19590, plain, ((e22 = op2(op2(e23, e25), e25)) | ~ spl42_424), inference(backward_demodulation, [], [f473, f3453])).
fof(f473, plain, (e22 = op2(op2(op2(e25, e22), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f19889, plain, (~ spl42_537 | ~ spl42_801), inference(avatar_split_clause, [], [f19369, f5449, f3919])).
fof(f3919, plain, (spl42_537 <=> (e24 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_537])])).
fof(f19369, plain, (~ (e24 = op2(e23, e20)) | ~ spl42_801), inference(forward_demodulation, [], [f4750, f5450])).
fof(f4750, plain, ~ (op2(e23, e20) = h4(e11)), inference(backward_demodulation, [], [f1009, f1231])).
fof(f1231, plain, (op2(e23, e23) = h4(e11)), inference(cnf_transformation, [], [f19])).
fof(f1009, plain, ~ (op2(e23, e20) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f19874, plain, (spl42_368 | ~ spl42_469 | ~ spl42_801), inference(avatar_split_clause, [], [f19571, f5449, f3637, f3219])).
fof(f19571, plain, ((e23 = op2(e26, e23)) | (~ spl42_469 | ~ spl42_801)), inference(backward_demodulation, [], [f19382, f3639])).
fof(f19382, plain, ((e23 = op2(op2(e24, e23), e23)) | ~ spl42_801), inference(forward_demodulation, [], [f4743, f5450])).
fof(f4743, plain, (e23 = op2(op2(h4(e11), e23), e23)), inference(backward_demodulation, [], [f478, f1231])).
fof(f478, plain, (e23 = op2(op2(op2(e23, e23), e23), e23)), inference(cnf_transformation, [], [f6])).
fof(f19865, plain, (spl42_340 | ~ spl42_288 | ~ spl42_317), inference(avatar_split_clause, [], [f19722, f2910, f2790, f3005])).
fof(f19722, plain, ((op1(e10, e10) = e13) | (~ spl42_288 | ~ spl42_317)), inference(forward_demodulation, [], [f19710, f2792])).
fof(f19710, plain, ((e13 = op1(op1(e11, e10), e10)) | ~ spl42_317), inference(backward_demodulation, [], [f279, f2912])).
fof(f279, plain, (e13 = op1(op1(op1(e10, e13), e10), e10)), inference(cnf_transformation, [], [f3])).
fof(f19850, plain, (spl42_287 | ~ spl42_233 | ~ spl42_248), inference(avatar_split_clause, [], [f19747, f2624, f2562, f2785])).
fof(f19747, plain, ((e16 = op1(e11, e11)) | (~ spl42_233 | ~ spl42_248)), inference(forward_demodulation, [], [f19741, f2564])).
fof(f19741, plain, ((e16 = op1(op1(e12, e11), e11)) | ~ spl42_248), inference(backward_demodulation, [], [f301, f2626])).
fof(f301, plain, (e16 = op1(op1(op1(e11, e16), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f19801, plain, (~ spl42_33 | ~ spl42_127 | ~ spl42_231), inference(avatar_contradiction_clause, [], [f19800])).
fof(f19800, plain, ($false | (~ spl42_33 | ~ spl42_127 | ~ spl42_231)), inference(subsumption_resolution, [], [f19799, f1092])).
fof(f1092, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f9])).
fof(f19799, plain, ((e10 = e12) | (~ spl42_33 | ~ spl42_127 | ~ spl42_231)), inference(forward_demodulation, [], [f19796, f2125])).
fof(f19796, plain, ((e12 = op1(e14, e12)) | (~ spl42_33 | ~ spl42_231)), inference(backward_demodulation, [], [f19760, f1735])).
fof(f19760, plain, ((e12 = op1(op1(e16, e12), e12)) | ~ spl42_231), inference(backward_demodulation, [], [f274, f2555])).
fof(f2555, plain, ((e16 = op1(e12, e12)) | ~ spl42_231), inference(avatar_component_clause, [], [f2553])).
fof(f2553, plain, (spl42_231 <=> (e16 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_231])])).
fof(f19798, plain, (~ spl42_26 | ~ spl42_33), inference(avatar_split_clause, [], [f19793, f1733, f1704])).
fof(f1704, plain, (spl42_26 <=> (e14 = op1(e16, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_26])])).
fof(f19793, plain, (~ (e14 = op1(e16, e13)) | ~ spl42_33), inference(backward_demodulation, [], [f787, f1735])).
fof(f787, plain, ~ (op1(e16, e12) = op1(e16, e13)), inference(cnf_transformation, [], [f7])).
fof(f19788, plain, (~ spl42_58 | ~ spl42_63), inference(avatar_contradiction_clause, [], [f19787])).
fof(f19787, plain, ($false | (~ spl42_58 | ~ spl42_63)), inference(subsumption_resolution, [], [f19786, f1101])).
fof(f1101, plain, ~ (e11 = e16), inference(cnf_transformation, [], [f9])).
fof(f19786, plain, ((e11 = e16) | (~ spl42_58 | ~ spl42_63)), inference(forward_demodulation, [], [f1859, f1839])).
fof(f1859, plain, ((e16 = op1(e15, e15)) | ~ spl42_63), inference(avatar_component_clause, [], [f1857])).
fof(f1857, plain, (spl42_63 <=> (e16 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_63])])).
fof(f19706, plain, (~ spl42_230 | ~ spl42_328), inference(avatar_split_clause, [], [f19696, f2955, f2549])).
fof(f2955, plain, (spl42_328 <=> (e15 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_328])])).
fof(f19696, plain, (~ (e15 = op1(e12, e12)) | ~ spl42_328), inference(backward_demodulation, [], [f546, f2957])).
fof(f2957, plain, ((e15 = op1(e10, e12)) | ~ spl42_328), inference(avatar_component_clause, [], [f2955])).
fof(f546, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f7])).
fof(f19606, plain, (~ spl42_423 | ~ spl42_424), inference(avatar_contradiction_clause, [], [f19605])).
fof(f19605, plain, ($false | (~ spl42_423 | ~ spl42_424)), inference(subsumption_resolution, [], [f19604, f1123])).
fof(f1123, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f10])).
fof(f19604, plain, ((e22 = e23) | (~ spl42_423 | ~ spl42_424)), inference(backward_demodulation, [], [f3453, f3449])).
fof(f3449, plain, ((e22 = op2(e25, e22)) | ~ spl42_423), inference(avatar_component_clause, [], [f3447])).
fof(f3447, plain, (spl42_423 <=> (e22 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_423])])).
fof(f19573, plain, (~ spl42_448 | ~ spl42_469), inference(avatar_split_clause, [], [f19569, f3637, f3550])).
fof(f3550, plain, (spl42_448 <=> (e26 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_448])])).
fof(f19569, plain, (~ (e26 = op2(e24, e26)) | ~ spl42_469), inference(backward_demodulation, [], [f1045, f3639])).
fof(f1045, plain, ~ (op2(e24, e23) = op2(e24, e26)), inference(cnf_transformation, [], [f8])).
fof(f19564, plain, (~ spl42_446 | ~ spl42_488), inference(avatar_split_clause, [], [f19558, f3716, f3542])).
fof(f3542, plain, (spl42_446 <=> (e24 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_446])])).
fof(f19558, plain, (~ (e24 = op2(e24, e26)) | ~ spl42_488), inference(backward_demodulation, [], [f1033, f3718])).
fof(f1033, plain, ~ (op2(e24, e20) = op2(e24, e26)), inference(cnf_transformation, [], [f8])).
fof(f19494, plain, (spl42_385 | ~ spl42_581 | ~ spl42_591), inference(avatar_contradiction_clause, [], [f19493])).
fof(f19493, plain, ($false | (spl42_385 | ~ spl42_581 | ~ spl42_591)), inference(subsumption_resolution, [], [f19492, f3290])).
fof(f3290, plain, (~ (e26 = op2(e26, e21)) | spl42_385), inference(avatar_component_clause, [], [f3289])).
fof(f3289, plain, (spl42_385 <=> (e26 = op2(e26, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_385])])).
fof(f19492, plain, ((e26 = op2(e26, e21)) | (~ spl42_581 | ~ spl42_591)), inference(forward_demodulation, [], [f19482, f4103])).
fof(f4103, plain, ((e26 = op2(e22, e21)) | ~ spl42_581), inference(avatar_component_clause, [], [f4101])).
fof(f4101, plain, (spl42_581 <=> (e26 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_581])])).
fof(f19482, plain, ((e26 = op2(op2(e22, e21), e21)) | ~ spl42_591), inference(backward_demodulation, [], [f497, f4145])).
fof(f497, plain, (e26 = op2(op2(op2(e21, e26), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f19379, plain, (spl42_1168 | ~ spl42_511 | ~ spl42_801), inference(avatar_split_clause, [], [f19378, f5449, f3811, f7615])).
fof(f3811, plain, (spl42_511 <=> (e26 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_511])])).
fof(f19378, plain, ((e26 = h4(e14)) | (~ spl42_511 | ~ spl42_801)), inference(forward_demodulation, [], [f19377, f3813])).
fof(f3813, plain, ((e26 = op2(e23, e24)) | ~ spl42_511), inference(avatar_component_clause, [], [f3811])).
fof(f19377, plain, ((op2(e23, e24) = h4(e14)) | ~ spl42_801), inference(forward_demodulation, [], [f4756, f5450])).
fof(f4756, plain, (h4(e14) = op2(e23, h4(e11))), inference(backward_demodulation, [], [f1234, f1231])).
fof(f19324, plain, (~ spl42_621 | ~ spl42_600), inference(avatar_split_clause, [], [f19323, f4180, f4267])).
fof(f4267, plain, (spl42_621 <=> (e24 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_621])])).
fof(f19323, plain, (~ (e24 = op2(e21, e22)) | ~ spl42_600), inference(forward_demodulation, [], [f978, f4182])).
fof(f978, plain, ~ (op2(e21, e22) = op2(e21, e25)), inference(cnf_transformation, [], [f8])).
fof(f19318, plain, (~ spl42_620 | ~ spl42_606), inference(avatar_split_clause, [], [f19317, f4205, f4263])).
fof(f4263, plain, (spl42_620 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_620])])).
fof(f19317, plain, (~ (e23 = op2(e21, e22)) | ~ spl42_606), inference(forward_demodulation, [], [f977, f4207])).
fof(f977, plain, ~ (op2(e21, e22) = op2(e21, e24)), inference(cnf_transformation, [], [f8])).
fof(f19308, plain, (spl42_463 | ~ spl42_516 | ~ spl42_536), inference(avatar_contradiction_clause, [], [f19307])).
fof(f19307, plain, ($false | (spl42_463 | ~ spl42_516 | ~ spl42_536)), inference(subsumption_resolution, [], [f19306, f3614])).
fof(f3614, plain, (~ (e20 = op2(e24, e23)) | spl42_463), inference(avatar_component_clause, [], [f3613])).
fof(f3613, plain, (spl42_463 <=> (e20 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_463])])).
fof(f19306, plain, ((e20 = op2(e24, e23)) | (~ spl42_516 | ~ spl42_536)), inference(forward_demodulation, [], [f19305, f3834])).
fof(f19305, plain, ((e20 = op2(op2(e23, e23), e23)) | ~ spl42_536), inference(forward_demodulation, [], [f457, f3917])).
fof(f3917, plain, ((e23 = op2(e23, e20)) | ~ spl42_536), inference(avatar_component_clause, [], [f3915])).
fof(f3915, plain, (spl42_536 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_536])])).
fof(f457, plain, (e20 = op2(op2(op2(e23, e20), e23), e23)), inference(cnf_transformation, [], [f6])).
fof(f19288, plain, (~ spl42_384 | ~ spl42_531), inference(avatar_split_clause, [], [f19286, f3894, f3285])).
fof(f3285, plain, (spl42_384 <=> (e25 = op2(e26, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_384])])).
fof(f19286, plain, (~ (e25 = op2(e26, e21)) | ~ spl42_531), inference(backward_demodulation, [], [f835, f3896])).
fof(f835, plain, ~ (op2(e23, e21) = op2(e26, e21)), inference(cnf_transformation, [], [f8])).
fof(f19285, plain, (spl42_531 | ~ spl42_480 | ~ spl42_600), inference(avatar_split_clause, [], [f18706, f4180, f3683, f3894])).
fof(f18706, plain, ((e25 = op2(e23, e21)) | (~ spl42_480 | ~ spl42_600)), inference(backward_demodulation, [], [f18375, f3685])).
fof(f18375, plain, ((e25 = op2(op2(e24, e21), e21)) | ~ spl42_600), inference(forward_demodulation, [], [f490, f4182])).
fof(f490, plain, (e25 = op2(op2(op2(e21, e25), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f19277, plain, (~ spl42_488 | ~ spl42_390), inference(avatar_split_clause, [], [f19276, f3310, f3716])).
fof(f3310, plain, (spl42_390 <=> (e24 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_390])])).
fof(f19276, plain, (~ (e24 = op2(e24, e20)) | ~ spl42_390), inference(forward_demodulation, [], [f816, f3312])).
fof(f3312, plain, ((e24 = op2(e26, e20)) | ~ spl42_390), inference(avatar_component_clause, [], [f3310])).
fof(f816, plain, ~ (op2(e24, e20) = op2(e26, e20)), inference(cnf_transformation, [], [f8])).
fof(f19272, plain, (~ spl42_468 | ~ spl42_454), inference(avatar_split_clause, [], [f19271, f3575, f3633])).
fof(f3633, plain, (spl42_468 <=> (e25 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_468])])).
fof(f19271, plain, (~ (e25 = op2(e24, e23)) | ~ spl42_454), inference(forward_demodulation, [], [f1044, f3577])).
fof(f1044, plain, ~ (op2(e24, e23) = op2(e24, e25)), inference(cnf_transformation, [], [f8])).
fof(f19269, plain, (~ spl42_469 | ~ spl42_616), inference(avatar_split_clause, [], [f19268, f4246, f3637])).
fof(f4246, plain, (spl42_616 <=> (e26 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_616])])).
fof(f19268, plain, (~ (e26 = op2(e24, e23)) | ~ spl42_616), inference(forward_demodulation, [], [f868, f4248])).
fof(f4248, plain, ((e26 = op2(e21, e23)) | ~ spl42_616), inference(avatar_component_clause, [], [f4246])).
fof(f868, plain, ~ (op2(e21, e23) = op2(e24, e23)), inference(cnf_transformation, [], [f8])).
fof(f19249, plain, (~ spl42_398 | ~ spl42_412), inference(avatar_split_clause, [], [f17601, f3401, f3343])).
fof(f3343, plain, (spl42_398 <=> (e25 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_398])])).
fof(f17601, plain, (~ (e25 = op2(e25, e26)) | ~ spl42_412), inference(backward_demodulation, [], [f1068, f3403])).
fof(f1068, plain, ~ (op2(e25, e24) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f19196, plain, (~ spl42_309 | ~ spl42_313), inference(avatar_contradiction_clause, [], [f19195])).
fof(f19195, plain, ($false | (~ spl42_309 | ~ spl42_313)), inference(subsumption_resolution, [], [f19194, f1094])).
fof(f1094, plain, ~ (e10 = e14), inference(cnf_transformation, [], [f9])).
fof(f19194, plain, ((e10 = e14) | (~ spl42_309 | ~ spl42_313)), inference(forward_demodulation, [], [f2895, f2879])).
fof(f2879, plain, ((e10 = op1(e10, e14)) | ~ spl42_309), inference(avatar_component_clause, [], [f2877])).
fof(f2877, plain, (spl42_309 <=> (e10 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_309])])).
fof(f19187, plain, (spl42_304 | ~ spl42_81 | ~ spl42_155), inference(avatar_split_clause, [], [f18989, f2239, f1932, f2856])).
fof(f18989, plain, ((e12 = op1(e10, e15)) | (~ spl42_81 | ~ spl42_155)), inference(backward_demodulation, [], [f17948, f2241])).
fof(f17948, plain, ((e12 = op1(op1(e13, e15), e15)) | ~ spl42_81), inference(backward_demodulation, [], [f277, f1934])).
fof(f277, plain, (e12 = op1(op1(op1(e15, e12), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f19166, plain, (~ spl42_271 | ~ spl42_257), inference(avatar_split_clause, [], [f19165, f2661, f2719])).
fof(f2719, plain, (spl42_271 <=> (e14 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_271])])).
fof(f19165, plain, (~ (e14 = op1(e11, e13)) | ~ spl42_257), inference(forward_demodulation, [], [f687, f2663])).
fof(f687, plain, ~ (op1(e11, e13) = op1(e11, e15)), inference(cnf_transformation, [], [f7])).
fof(f19136, plain, (spl42_230 | ~ spl42_178 | ~ spl42_207), inference(avatar_split_clause, [], [f18966, f2454, f2334, f2549])).
fof(f18966, plain, ((e15 = op1(e12, e12)) | (~ spl42_178 | ~ spl42_207)), inference(backward_demodulation, [], [f13407, f2336])).
fof(f19098, plain, (~ spl42_55 | ~ spl42_69), inference(avatar_split_clause, [], [f19097, f1882, f1824])).
fof(f1824, plain, (spl42_55 <=> (e15 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_55])])).
fof(f19097, plain, (~ (e15 = op1(e15, e16)) | ~ spl42_69), inference(forward_demodulation, [], [f774, f1884])).
fof(f774, plain, ~ (op1(e15, e14) = op1(e15, e16)), inference(cnf_transformation, [], [f7])).
fof(f19002, plain, (spl42_126 | ~ spl42_151 | ~ spl42_173), inference(avatar_split_clause, [], [f19001, f2313, f2222, f2118])).
fof(f19001, plain, ((e16 = op1(e14, e13)) | (~ spl42_151 | ~ spl42_173)), inference(forward_demodulation, [], [f18996, f2315])).
fof(f18996, plain, ((e16 = op1(op1(e13, e13), e13)) | ~ spl42_151), inference(backward_demodulation, [], [f303, f2224])).
fof(f303, plain, (e16 = op1(op1(op1(e13, e16), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f18993, plain, (~ spl42_155 | ~ spl42_173 | ~ spl42_319), inference(avatar_contradiction_clause, [], [f18992])).
fof(f18992, plain, ($false | (~ spl42_155 | ~ spl42_173 | ~ spl42_319)), inference(subsumption_resolution, [], [f18991, f1109])).
fof(f1109, plain, ~ (e14 = e15), inference(cnf_transformation, [], [f9])).
fof(f18991, plain, ((e14 = e15) | (~ spl42_155 | ~ spl42_173 | ~ spl42_319)), inference(forward_demodulation, [], [f18990, f2315])).
fof(f18990, plain, ((e15 = op1(e13, e13)) | (~ spl42_155 | ~ spl42_319)), inference(forward_demodulation, [], [f18980, f2920])).
fof(f2920, plain, ((e13 = op1(e10, e13)) | ~ spl42_319), inference(avatar_component_clause, [], [f2918])).
fof(f2918, plain, (spl42_319 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_319])])).
fof(f18980, plain, ((e15 = op1(op1(e10, e13), e13)) | ~ spl42_155), inference(backward_demodulation, [], [f296, f2241])).
fof(f296, plain, (e15 = op1(op1(op1(e13, e15), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f18973, plain, (~ spl42_178 | ~ spl42_207 | ~ spl42_229), inference(avatar_contradiction_clause, [], [f18972])).
fof(f18972, plain, ($false | (~ spl42_178 | ~ spl42_207 | ~ spl42_229)), inference(subsumption_resolution, [], [f18971, f1109])).
fof(f18971, plain, ((e14 = e15) | (~ spl42_178 | ~ spl42_207 | ~ spl42_229)), inference(forward_demodulation, [], [f18966, f2547])).
fof(f2547, plain, ((e14 = op1(e12, e12)) | ~ spl42_229), inference(avatar_component_clause, [], [f2545])).
fof(f2545, plain, (spl42_229 <=> (e14 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_229])])).
fof(f18957, plain, (~ spl42_185 | ~ spl42_188), inference(avatar_contradiction_clause, [], [f18956])).
fof(f18956, plain, ($false | (~ spl42_185 | ~ spl42_188)), inference(subsumption_resolution, [], [f18955, f1104])).
fof(f1104, plain, ~ (e12 = e15), inference(cnf_transformation, [], [f9])).
fof(f18955, plain, ((e12 = e15) | (~ spl42_185 | ~ spl42_188)), inference(backward_demodulation, [], [f2377, f2365])).
fof(f2365, plain, ((e12 = op1(e13, e11)) | ~ spl42_185), inference(avatar_component_clause, [], [f2363])).
fof(f2363, plain, (spl42_185 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_185])])).
fof(f18923, plain, (~ spl42_207 | ~ spl42_210), inference(avatar_contradiction_clause, [], [f18922])).
fof(f18922, plain, ($false | (~ spl42_207 | ~ spl42_210)), inference(subsumption_resolution, [], [f18921, f1108])).
fof(f1108, plain, ~ (e13 = e16), inference(cnf_transformation, [], [f9])).
fof(f18921, plain, ((e13 = e16) | (~ spl42_207 | ~ spl42_210)), inference(forward_demodulation, [], [f2468, f2456])).
fof(f2468, plain, ((e16 = op1(e12, e15)) | ~ spl42_210), inference(avatar_component_clause, [], [f2466])).
fof(f2466, plain, (spl42_210 <=> (e16 = op1(e12, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_210])])).
fof(f18770, plain, (~ spl42_390 | spl42_442 | ~ spl42_446), inference(avatar_contradiction_clause, [], [f18769])).
fof(f18769, plain, ($false | (~ spl42_390 | spl42_442 | ~ spl42_446)), inference(subsumption_resolution, [], [f18768, f3527])).
fof(f3527, plain, (~ (e20 = op2(e24, e26)) | spl42_442), inference(avatar_component_clause, [], [f3526])).
fof(f3526, plain, (spl42_442 <=> (e20 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_442])])).
fof(f18768, plain, ((e20 = op2(e24, e26)) | (~ spl42_390 | ~ spl42_446)), inference(forward_demodulation, [], [f18764, f3544])).
fof(f3544, plain, ((e24 = op2(e24, e26)) | ~ spl42_446), inference(avatar_component_clause, [], [f3542])).
fof(f18764, plain, ((e20 = op2(op2(e24, e26), e26)) | ~ spl42_390), inference(backward_demodulation, [], [f460, f3312])).
fof(f18730, plain, (~ spl42_435 | ~ spl42_437), inference(avatar_contradiction_clause, [], [f18729])).
fof(f18729, plain, ($false | (~ spl42_435 | ~ spl42_437)), inference(subsumption_resolution, [], [f18728, f1113])).
fof(f1113, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f10])).
fof(f18728, plain, ((e20 = e22) | (~ spl42_435 | ~ spl42_437)), inference(backward_demodulation, [], [f3507, f3499])).
fof(f3499, plain, ((e20 = op2(e25, e20)) | ~ spl42_435), inference(avatar_component_clause, [], [f3497])).
fof(f3497, plain, (spl42_435 <=> (e20 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_435])])).
fof(f18711, plain, (~ spl42_480 | ~ spl42_528 | ~ spl42_600), inference(avatar_contradiction_clause, [], [f18710])).
fof(f18710, plain, ($false | (~ spl42_480 | ~ spl42_528 | ~ spl42_600)), inference(subsumption_resolution, [], [f18709, f1125])).
fof(f1125, plain, ~ (e22 = e25), inference(cnf_transformation, [], [f10])).
fof(f18709, plain, ((e22 = e25) | (~ spl42_480 | ~ spl42_528 | ~ spl42_600)), inference(forward_demodulation, [], [f18706, f3884])).
fof(f3884, plain, ((e22 = op2(e23, e21)) | ~ spl42_528), inference(avatar_component_clause, [], [f3882])).
fof(f3882, plain, (spl42_528 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_528])])).
fof(f18700, plain, (spl42_390 | ~ spl42_490 | ~ spl42_656), inference(avatar_split_clause, [], [f18695, f4412, f3724, f3310])).
fof(f3724, plain, (spl42_490 <=> (e26 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_490])])).
fof(f18695, plain, ((e24 = op2(e26, e20)) | (~ spl42_490 | ~ spl42_656)), inference(backward_demodulation, [], [f8815, f3726])).
fof(f3726, plain, ((e26 = op2(e24, e20)) | ~ spl42_490), inference(avatar_component_clause, [], [f3724])).
fof(f8815, plain, ((e24 = op2(op2(e24, e20), e20)) | ~ spl42_656), inference(backward_demodulation, [], [f482, f4414])).
fof(f482, plain, (e24 = op2(op2(op2(e20, e24), e20), e20)), inference(cnf_transformation, [], [f6])).
fof(f18684, plain, (~ spl42_522 | ~ spl42_550), inference(avatar_contradiction_clause, [], [f18683])).
fof(f18683, plain, ($false | (~ spl42_522 | ~ spl42_550)), inference(subsumption_resolution, [], [f18682, f1128])).
fof(f1128, plain, ~ (e23 = e25), inference(cnf_transformation, [], [f10])).
fof(f18682, plain, ((e23 = e25) | (~ spl42_522 | ~ spl42_550)), inference(forward_demodulation, [], [f18676, f3859])).
fof(f3859, plain, ((e23 = op2(e23, e22)) | ~ spl42_522), inference(avatar_component_clause, [], [f3857])).
fof(f3857, plain, (spl42_522 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_522])])).
fof(f18676, plain, ((e25 = op2(e23, e22)) | (~ spl42_522 | ~ spl42_550)), inference(backward_demodulation, [], [f13666, f3859])).
fof(f18634, plain, (~ spl42_593 | ~ spl42_600), inference(avatar_contradiction_clause, [], [f18633])).
fof(f18633, plain, ($false | (~ spl42_593 | ~ spl42_600)), inference(subsumption_resolution, [], [f18632, f1132])).
fof(f1132, plain, ~ (e25 = e26), inference(cnf_transformation, [], [f10])).
fof(f18632, plain, ((e25 = e26) | (~ spl42_593 | ~ spl42_600)), inference(forward_demodulation, [], [f18627, f18375])).
fof(f18627, plain, ((e26 = op2(op2(e24, e21), e21)) | ~ spl42_593), inference(backward_demodulation, [], [f497, f4153])).
fof(f4153, plain, ((e24 = op2(e21, e26)) | ~ spl42_593), inference(avatar_component_clause, [], [f4151])).
fof(f4151, plain, (spl42_593 <=> (e24 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_593])])).
fof(f18614, plain, (~ spl42_638 | ~ spl42_643), inference(avatar_contradiction_clause, [], [f18613])).
fof(f18613, plain, ($false | (~ spl42_638 | ~ spl42_643)), inference(subsumption_resolution, [], [f18612, f1116])).
fof(f1116, plain, ~ (e20 = e25), inference(cnf_transformation, [], [f10])).
fof(f18612, plain, ((e20 = e25) | (~ spl42_638 | ~ spl42_643)), inference(forward_demodulation, [], [f4360, f4340])).
fof(f4340, plain, ((e20 = op2(e20, e26)) | ~ spl42_638), inference(avatar_component_clause, [], [f4338])).
fof(f4338, plain, (spl42_638 <=> (e20 = op2(e20, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_638])])).
fof(f18603, plain, (~ spl42_437 | spl42_585 | ~ spl42_664), inference(avatar_contradiction_clause, [], [f18602])).
fof(f18602, plain, ($false | (~ spl42_437 | spl42_585 | ~ spl42_664)), inference(subsumption_resolution, [], [f18601, f4119])).
fof(f4119, plain, (~ (e23 = op2(e22, e20)) | spl42_585), inference(avatar_component_clause, [], [f4118])).
fof(f4118, plain, (spl42_585 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_585])])).
fof(f18601, plain, ((e23 = op2(e22, e20)) | (~ spl42_437 | ~ spl42_664)), inference(forward_demodulation, [], [f18596, f3507])).
fof(f18596, plain, ((e23 = op2(op2(e25, e20), e20)) | ~ spl42_664), inference(backward_demodulation, [], [f475, f4447])).
fof(f4447, plain, ((e25 = op2(e20, e23)) | ~ spl42_664), inference(avatar_component_clause, [], [f4445])).
fof(f4445, plain, (spl42_664 <=> (e25 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_664])])).
fof(f475, plain, (e23 = op2(op2(op2(e20, e23), e20), e20)), inference(cnf_transformation, [], [f6])).
fof(f18386, plain, (~ spl42_614 | ~ spl42_600), inference(avatar_split_clause, [], [f18385, f4180, f4238])).
fof(f4238, plain, (spl42_614 <=> (e24 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_614])])).
fof(f18385, plain, (~ (e24 = op2(e21, e23)) | ~ spl42_600), inference(forward_demodulation, [], [f981, f4182])).
fof(f981, plain, ~ (op2(e21, e23) = op2(e21, e25)), inference(cnf_transformation, [], [f8])).
fof(f18230, plain, (~ spl42_329 | ~ spl42_308), inference(avatar_split_clause, [], [f14397, f2872, f2959])).
fof(f2872, plain, (spl42_308 <=> (e16 = op1(e10, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_308])])).
fof(f14397, plain, (~ (e16 = op1(e10, e12)) | ~ spl42_308), inference(forward_demodulation, [], [f663, f2874])).
fof(f2874, plain, ((e16 = op1(e10, e15)) | ~ spl42_308), inference(avatar_component_clause, [], [f2872])).
fof(f663, plain, ~ (op1(e10, e12) = op1(e10, e15)), inference(cnf_transformation, [], [f7])).
fof(f18219, plain, (spl42_320 | ~ spl42_22 | ~ spl42_168), inference(avatar_split_clause, [], [f18028, f2292, f1688, f2922])).
fof(f2922, plain, (spl42_320 <=> (e14 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_320])])).
fof(f1688, plain, (spl42_22 <=> (e10 = op1(e16, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_22])])).
fof(f18028, plain, ((e14 = op1(e10, e13)) | (~ spl42_22 | ~ spl42_168)), inference(backward_demodulation, [], [f10201, f1690])).
fof(f1690, plain, ((e10 = op1(e16, e13)) | ~ spl42_22), inference(avatar_component_clause, [], [f1688])).
fof(f10201, plain, ((e14 = op1(op1(e16, e13), e13)) | ~ spl42_168), inference(backward_demodulation, [], [f289, f2294])).
fof(f289, plain, (e14 = op1(op1(op1(e13, e14), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f18218, plain, (spl42_14 | ~ spl42_50 | ~ spl42_308), inference(avatar_contradiction_clause, [], [f18217])).
fof(f18217, plain, ($false | (spl42_14 | ~ spl42_50 | ~ spl42_308)), inference(subsumption_resolution, [], [f18216, f1655])).
fof(f1655, plain, (~ (e16 = op1(e16, e15)) | spl42_14), inference(avatar_component_clause, [], [f1654])).
fof(f18216, plain, ((e16 = op1(e16, e15)) | (~ spl42_50 | ~ spl42_308)), inference(backward_demodulation, [], [f18113, f2874])).
fof(f18113, plain, ((e16 = op1(op1(e10, e15), e15)) | ~ spl42_50), inference(forward_demodulation, [], [f305, f1806])).
fof(f1806, plain, ((e10 = op1(e15, e16)) | ~ spl42_50), inference(avatar_component_clause, [], [f1804])).
fof(f1804, plain, (spl42_50 <=> (e10 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_50])])).
fof(f305, plain, (e16 = op1(op1(op1(e15, e16), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f18175, plain, (spl42_326 | ~ spl42_127 | ~ spl42_222), inference(avatar_split_clause, [], [f18174, f2516, f2123, f2947])).
fof(f2947, plain, (spl42_326 <=> (e13 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_326])])).
fof(f2516, plain, (spl42_222 <=> (e14 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_222])])).
fof(f18174, plain, ((e13 = op1(e10, e12)) | (~ spl42_127 | ~ spl42_222)), inference(forward_demodulation, [], [f18173, f2125])).
fof(f18173, plain, ((e13 = op1(op1(e14, e12), e12)) | ~ spl42_222), inference(forward_demodulation, [], [f281, f2518])).
fof(f2518, plain, ((e14 = op1(e12, e13)) | ~ spl42_222), inference(avatar_component_clause, [], [f2516])).
fof(f18055, plain, (~ spl42_48 | ~ spl42_6), inference(avatar_split_clause, [], [f18054, f1621, f1795])).
fof(f1621, plain, (spl42_6 <=> (e15 = op1(e16, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_6])])).
fof(f18054, plain, (~ (e15 = op1(e16, e10)) | ~ spl42_6), inference(forward_demodulation, [], [f781, f1623])).
fof(f1623, plain, ((e15 = op1(e16, e16)) | ~ spl42_6), inference(avatar_component_clause, [], [f1621])).
fof(f781, plain, ~ (op1(e16, e10) = op1(e16, e16)), inference(cnf_transformation, [], [f7])).
fof(f18015, plain, (~ spl42_33 | ~ spl42_127 | ~ spl42_238), inference(avatar_contradiction_clause, [], [f18014])).
fof(f18014, plain, ($false | (~ spl42_33 | ~ spl42_127 | ~ spl42_238)), inference(subsumption_resolution, [], [f18013, f1091])).
fof(f1091, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f9])).
fof(f18013, plain, ((e10 = e11) | (~ spl42_33 | ~ spl42_127 | ~ spl42_238)), inference(forward_demodulation, [], [f18009, f2125])).
fof(f18009, plain, ((e11 = op1(e14, e12)) | (~ spl42_33 | ~ spl42_238)), inference(backward_demodulation, [], [f17838, f1735])).
fof(f17838, plain, ((e11 = op1(op1(e16, e12), e12)) | ~ spl42_238), inference(backward_demodulation, [], [f267, f2584])).
fof(f2584, plain, ((e16 = op1(e12, e11)) | ~ spl42_238), inference(avatar_component_clause, [], [f2582])).
fof(f2582, plain, (spl42_238 <=> (e16 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_238])])).
fof(f17979, plain, (~ spl42_67 | ~ spl42_111 | spl42_167), inference(avatar_contradiction_clause, [], [f17978])).
fof(f17978, plain, ($false | (~ spl42_67 | ~ spl42_111 | spl42_167)), inference(subsumption_resolution, [], [f17973, f2289])).
fof(f2289, plain, (~ (e15 = op1(e13, e14)) | spl42_167), inference(avatar_component_clause, [], [f2288])).
fof(f2288, plain, (spl42_167 <=> (e15 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_167])])).
fof(f17973, plain, ((e15 = op1(e13, e14)) | (~ spl42_67 | ~ spl42_111)), inference(backward_demodulation, [], [f15056, f1876])).
fof(f1876, plain, ((e13 = op1(e15, e14)) | ~ spl42_67), inference(avatar_component_clause, [], [f1874])).
fof(f1874, plain, (spl42_67 <=> (e13 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_67])])).
fof(f15056, plain, ((e15 = op1(op1(e15, e14), e14)) | ~ spl42_111), inference(backward_demodulation, [], [f297, f2058])).
fof(f297, plain, (e15 = op1(op1(op1(e14, e15), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f17929, plain, (~ spl42_16 | ~ spl42_123 | ~ spl42_168), inference(avatar_contradiction_clause, [], [f17928])).
fof(f17928, plain, ($false | (~ spl42_16 | ~ spl42_123 | ~ spl42_168)), inference(subsumption_resolution, [], [f17927, f1098])).
fof(f1098, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f9])).
fof(f17927, plain, ((e11 = e13) | (~ spl42_16 | ~ spl42_123 | ~ spl42_168)), inference(forward_demodulation, [], [f17926, f1665])).
fof(f17926, plain, ((e13 = op1(e16, e14)) | (~ spl42_123 | ~ spl42_168)), inference(forward_demodulation, [], [f17919, f2294])).
fof(f17919, plain, ((e13 = op1(op1(e13, e14), e14)) | ~ spl42_123), inference(backward_demodulation, [], [f283, f2108])).
fof(f2108, plain, ((e13 = op1(e14, e13)) | ~ spl42_123), inference(avatar_component_clause, [], [f2106])).
fof(f2106, plain, (spl42_123 <=> (e13 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_123])])).
fof(f283, plain, (e13 = op1(op1(op1(e14, e13), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f17871, plain, (~ spl42_179 | ~ spl42_207), inference(avatar_contradiction_clause, [], [f17870])).
fof(f17870, plain, ($false | (~ spl42_179 | ~ spl42_207)), inference(subsumption_resolution, [], [f17869, f1107])).
fof(f1107, plain, ~ (e13 = e15), inference(cnf_transformation, [], [f9])).
fof(f17869, plain, ((e13 = e15) | (~ spl42_179 | ~ spl42_207)), inference(forward_demodulation, [], [f17862, f2340])).
fof(f2340, plain, ((e13 = op1(e13, e12)) | ~ spl42_179), inference(avatar_component_clause, [], [f2338])).
fof(f2338, plain, (spl42_179 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_179])])).
fof(f17862, plain, ((e15 = op1(e13, e12)) | (~ spl42_179 | ~ spl42_207)), inference(backward_demodulation, [], [f13407, f2340])).
fof(f17746, plain, (~ spl42_89 | spl42_134 | ~ spl42_293), inference(avatar_contradiction_clause, [], [f17745])).
fof(f17745, plain, ($false | (~ spl42_89 | spl42_134 | ~ spl42_293)), inference(subsumption_resolution, [], [f17744, f2153])).
fof(f2153, plain, (~ (e10 = op1(e14, e11)) | spl42_134), inference(avatar_component_clause, [], [f2152])).
fof(f2152, plain, (spl42_134 <=> (e10 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_134])])).
fof(f17744, plain, ((e10 = op1(e14, e11)) | (~ spl42_89 | ~ spl42_293)), inference(forward_demodulation, [], [f17734, f1967])).
fof(f17734, plain, ((e10 = op1(op1(e15, e11), e11)) | ~ spl42_293), inference(backward_demodulation, [], [f259, f2812])).
fof(f2812, plain, ((e15 = op1(e11, e10)) | ~ spl42_293), inference(avatar_component_clause, [], [f2810])).
fof(f2810, plain, (spl42_293 <=> (e15 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_293])])).
fof(f259, plain, (e10 = op1(op1(op1(e11, e10), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f17705, plain, (~ spl42_211 | spl42_278 | ~ spl42_324), inference(avatar_contradiction_clause, [], [f17704])).
fof(f17704, plain, ($false | (~ spl42_211 | spl42_278 | ~ spl42_324)), inference(subsumption_resolution, [], [f17698, f2749])).
fof(f2749, plain, (~ (e14 = op1(e11, e12)) | spl42_278), inference(avatar_component_clause, [], [f2748])).
fof(f17698, plain, ((e14 = op1(e11, e12)) | (~ spl42_211 | ~ spl42_324)), inference(backward_demodulation, [], [f13086, f2941])).
fof(f2941, plain, ((e11 = op1(e10, e12)) | ~ spl42_324), inference(avatar_component_clause, [], [f2939])).
fof(f2939, plain, (spl42_324 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_324])])).
fof(f13086, plain, ((e14 = op1(op1(e10, e12), e12)) | ~ spl42_211), inference(backward_demodulation, [], [f288, f2473])).
fof(f288, plain, (e14 = op1(op1(op1(e12, e14), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f17622, plain, (~ spl42_393 | ~ spl42_550 | ~ spl42_647), inference(avatar_contradiction_clause, [], [f17621])).
fof(f17621, plain, ($false | (~ spl42_393 | ~ spl42_550 | ~ spl42_647)), inference(subsumption_resolution, [], [f17620, f1129])).
fof(f1129, plain, ~ (e23 = e26), inference(cnf_transformation, [], [f10])).
fof(f17620, plain, ((e23 = e26) | (~ spl42_393 | ~ spl42_550 | ~ spl42_647)), inference(forward_demodulation, [], [f17619, f3975])).
fof(f17619, plain, ((e26 = op2(e22, e25)) | (~ spl42_393 | ~ spl42_647)), inference(forward_demodulation, [], [f17610, f4377])).
fof(f17610, plain, ((e26 = op2(op2(e20, e25), e25)) | ~ spl42_393), inference(backward_demodulation, [], [f501, f3325])).
fof(f3325, plain, ((e20 = op2(e25, e26)) | ~ spl42_393), inference(avatar_component_clause, [], [f3323])).
fof(f501, plain, (e26 = op2(op2(op2(e25, e26), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f17585, plain, (~ spl42_437 | ~ spl42_441), inference(avatar_contradiction_clause, [], [f17584])).
fof(f17584, plain, ($false | (~ spl42_437 | ~ spl42_441)), inference(subsumption_resolution, [], [f17583, f1126])).
fof(f1126, plain, ~ (e22 = e26), inference(cnf_transformation, [], [f10])).
fof(f17583, plain, ((e22 = e26) | (~ spl42_437 | ~ spl42_441)), inference(backward_demodulation, [], [f3523, f3507])).
fof(f3523, plain, ((e26 = op2(e25, e20)) | ~ spl42_441), inference(avatar_component_clause, [], [f3521])).
fof(f3521, plain, (spl42_441 <=> (e26 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_441])])).
fof(f17562, plain, (spl42_352 | ~ spl42_432 | ~ spl42_455), inference(avatar_contradiction_clause, [], [f17561])).
fof(f17561, plain, ($false | (spl42_352 | ~ spl42_432 | ~ spl42_455)), inference(subsumption_resolution, [], [f17559, f3154])).
fof(f3154, plain, (~ (e21 = op2(e26, e25)) | spl42_352), inference(avatar_component_clause, [], [f3153])).
fof(f3153, plain, (spl42_352 <=> (e21 = op2(e26, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_352])])).
fof(f17559, plain, ((e21 = op2(e26, e25)) | (~ spl42_432 | ~ spl42_455)), inference(backward_demodulation, [], [f12593, f3581])).
fof(f3581, plain, ((e26 = op2(e24, e25)) | ~ spl42_455), inference(avatar_component_clause, [], [f3579])).
fof(f3579, plain, (spl42_455 <=> (e26 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_455])])).
fof(f12593, plain, ((e21 = op2(op2(e24, e25), e25)) | ~ spl42_432), inference(forward_demodulation, [], [f466, f3486])).
fof(f466, plain, (e21 = op2(op2(op2(e25, e21), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f17556, plain, (~ spl42_470 | ~ spl42_473), inference(avatar_contradiction_clause, [], [f17555])).
fof(f17555, plain, ($false | (~ spl42_470 | ~ spl42_473)), inference(subsumption_resolution, [], [f17554, f1114])).
fof(f1114, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f10])).
fof(f17554, plain, ((e20 = e23) | (~ spl42_470 | ~ spl42_473)), inference(forward_demodulation, [], [f3656, f3644])).
fof(f3656, plain, ((e23 = op2(e24, e22)) | ~ spl42_473), inference(avatar_component_clause, [], [f3654])).
fof(f3654, plain, (spl42_473 <=> (e23 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_473])])).
fof(f17553, plain, (~ spl42_470 | ~ spl42_474), inference(avatar_contradiction_clause, [], [f17552])).
fof(f17552, plain, ($false | (~ spl42_470 | ~ spl42_474)), inference(subsumption_resolution, [], [f17551, f1115])).
fof(f17551, plain, ((e20 = e24) | (~ spl42_470 | ~ spl42_474)), inference(forward_demodulation, [], [f3660, f3644])).
fof(f3660, plain, ((e24 = op2(e24, e22)) | ~ spl42_474), inference(avatar_component_clause, [], [f3658])).
fof(f3658, plain, (spl42_474 <=> (e24 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_474])])).
fof(f17509, plain, (~ spl42_470 | ~ spl42_565 | spl42_669), inference(avatar_contradiction_clause, [], [f17508])).
fof(f17508, plain, ($false | (~ spl42_470 | ~ spl42_565 | spl42_669)), inference(subsumption_resolution, [], [f17507, f4467])).
fof(f4467, plain, (~ (e23 = op2(e20, e22)) | spl42_669), inference(avatar_component_clause, [], [f4466])).
fof(f4466, plain, (spl42_669 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_669])])).
fof(f17507, plain, ((e23 = op2(e20, e22)) | (~ spl42_470 | ~ spl42_565)), inference(forward_demodulation, [], [f17501, f3644])).
fof(f17501, plain, ((e23 = op2(op2(e24, e22), e22)) | ~ spl42_565), inference(backward_demodulation, [], [f477, f4037])).
fof(f4037, plain, ((e24 = op2(e22, e23)) | ~ spl42_565), inference(avatar_component_clause, [], [f4035])).
fof(f4035, plain, (spl42_565 <=> (e24 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_565])])).
fof(f477, plain, (e23 = op2(op2(op2(e22, e23), e22), e22)), inference(cnf_transformation, [], [f6])).
fof(f17477, plain, (~ spl42_592 | ~ spl42_606), inference(avatar_split_clause, [], [f17467, f4205, f4147])).
fof(f4147, plain, (spl42_592 <=> (e23 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_592])])).
fof(f17467, plain, (~ (e23 = op2(e21, e26)) | ~ spl42_606), inference(backward_demodulation, [], [f984, f4207])).
fof(f984, plain, ~ (op2(e21, e24) = op2(e21, e26)), inference(cnf_transformation, [], [f8])).
fof(f17475, plain, (~ spl42_529 | ~ spl42_606), inference(avatar_contradiction_clause, [], [f17474])).
fof(f17474, plain, ($false | (~ spl42_529 | ~ spl42_606)), inference(subsumption_resolution, [], [f17473, f1127])).
fof(f1127, plain, ~ (e23 = e24), inference(cnf_transformation, [], [f10])).
fof(f17473, plain, ((e23 = e24) | (~ spl42_529 | ~ spl42_606)), inference(forward_demodulation, [], [f17472, f3888])).
fof(f3888, plain, ((e23 = op2(e23, e21)) | ~ spl42_529), inference(avatar_component_clause, [], [f3886])).
fof(f3886, plain, (spl42_529 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_529])])).
fof(f17472, plain, ((e24 = op2(e23, e21)) | (~ spl42_529 | ~ spl42_606)), inference(forward_demodulation, [], [f17464, f3888])).
fof(f17464, plain, ((e24 = op2(op2(e23, e21), e21)) | ~ spl42_606), inference(backward_demodulation, [], [f483, f4207])).
fof(f483, plain, (e24 = op2(op2(op2(e21, e24), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f17444, plain, (~ spl42_606 | ~ spl42_634), inference(avatar_split_clause, [], [f17434, f4321, f4205])).
fof(f4321, plain, (spl42_634 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_634])])).
fof(f17434, plain, (~ (e23 = op2(e21, e24)) | ~ spl42_634), inference(backward_demodulation, [], [f968, f4323])).
fof(f4323, plain, ((e23 = op2(e21, e20)) | ~ spl42_634), inference(avatar_component_clause, [], [f4321])).
fof(f968, plain, ~ (op2(e21, e20) = op2(e21, e24)), inference(cnf_transformation, [], [f8])).
fof(f17310, plain, (spl42_613 | ~ spl42_529), inference(avatar_split_clause, [], [f17309, f3886, f4234])).
fof(f4234, plain, (spl42_613 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_613])])).
fof(f17309, plain, ((e23 = op2(e21, e23)) | ~ spl42_529), inference(backward_demodulation, [], [f4743, f17308])).
fof(f17308, plain, ((e21 = op2(h4(e11), e23)) | ~ spl42_529), inference(backward_demodulation, [], [f16437, f1231])).
fof(f16437, plain, ((e21 = op2(op2(e23, e23), e23)) | ~ spl42_529), inference(backward_demodulation, [], [f464, f3888])).
fof(f17276, plain, (~ spl42_672 | ~ spl42_1336), inference(avatar_split_clause, [], [f17275, f8710, f4478])).
fof(f8710, plain, (spl42_1336 <=> (e26 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_1336])])).
fof(f17275, plain, (~ (e26 = op2(e20, e22)) | ~ spl42_1336), inference(forward_demodulation, [], [f4685, f8711])).
fof(f8711, plain, ((e26 = h1(e11)) | ~ spl42_1336), inference(avatar_component_clause, [], [f8710])).
fof(f4685, plain, ~ (op2(e20, e22) = h1(e11)), inference(backward_demodulation, [], [f945, f1210])).
fof(f1210, plain, (op2(e20, e20) = h1(e11)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((h1(e16) = op2(op2(op2(op2(e20, op2(e20, e20)), op2(e20, op2(e20, e20))), e20), op2(e20, op2(e20, e20)))) & (op2(e20, op2(e20, e20)) = h1(e14)) & (h1(e13) = op2(op2(op2(e20, op2(e20, e20)), op2(e20, op2(e20, e20))), e20)) & (op2(op2(e20, op2(e20, e20)), op2(e20, op2(e20, e20))) = h1(e12)) & (op2(e20, e20) = h1(e11)) & (h1(e10) = op2(op2(op2(e20, op2(e20, e20)), op2(e20, op2(e20, e20))), op2(e20, op2(e20, e20)))) & (e20 = h1(e15))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax16)).
fof(f945, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f8])).
fof(f17205, plain, (~ spl42_566 | ~ spl42_836), inference(avatar_split_clause, [], [f17204, f5625, f4039])).
fof(f4039, plain, (spl42_566 <=> (e25 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_566])])).
fof(f5625, plain, (spl42_836 <=> (e25 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_836])])).
fof(f17204, plain, (~ (e25 = op2(e22, e23)) | ~ spl42_836), inference(forward_demodulation, [], [f4730, f5626])).
fof(f5626, plain, ((e25 = h3(e11)) | ~ spl42_836), inference(avatar_component_clause, [], [f5625])).
fof(f4730, plain, ~ (op2(e22, e23) = h3(e11)), inference(backward_demodulation, [], [f997, f1224])).
fof(f997, plain, ~ (op2(e22, e22) = op2(e22, e23)), inference(cnf_transformation, [], [f8])).
fof(f17134, plain, (~ spl42_470 | ~ spl42_586 | spl42_666), inference(avatar_contradiction_clause, [], [f17133])).
fof(f17133, plain, ($false | (~ spl42_470 | ~ spl42_586 | spl42_666)), inference(subsumption_resolution, [], [f17132, f4455])).
fof(f4455, plain, (~ (e20 = op2(e20, e22)) | spl42_666), inference(avatar_component_clause, [], [f4454])).
fof(f4454, plain, (spl42_666 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_666])])).
fof(f17132, plain, ((e20 = op2(e20, e22)) | (~ spl42_470 | ~ spl42_586)), inference(forward_demodulation, [], [f16398, f3644])).
fof(f16398, plain, ((e20 = op2(op2(e24, e22), e22)) | ~ spl42_586), inference(backward_demodulation, [], [f456, f4124])).
fof(f4124, plain, ((e24 = op2(e22, e20)) | ~ spl42_586), inference(avatar_component_clause, [], [f4122])).
fof(f4122, plain, (spl42_586 <=> (e24 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_586])])).
fof(f456, plain, (e20 = op2(op2(op2(e22, e20), e22), e22)), inference(cnf_transformation, [], [f6])).
fof(f17126, plain, (~ spl42_425 | ~ spl42_432), inference(avatar_split_clause, [], [f17125, f3484, f3455])).
fof(f3455, plain, (spl42_425 <=> (e24 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_425])])).
fof(f17125, plain, (~ (e24 = op2(e25, e22)) | ~ spl42_432), inference(forward_demodulation, [], [f1055, f3486])).
fof(f1055, plain, ~ (op2(e25, e21) = op2(e25, e22)), inference(cnf_transformation, [], [f8])).
fof(f17043, plain, (~ spl42_58 | ~ spl42_207 | ~ spl42_255), inference(avatar_contradiction_clause, [], [f17042])).
fof(f17042, plain, ($false | (~ spl42_58 | ~ spl42_207 | ~ spl42_255)), inference(subsumption_resolution, [], [f17041, f1107])).
fof(f17041, plain, ((e13 = e15) | (~ spl42_58 | ~ spl42_207 | ~ spl42_255)), inference(forward_demodulation, [], [f17040, f2456])).
fof(f17040, plain, ((e15 = op1(e12, e15)) | (~ spl42_58 | ~ spl42_255)), inference(forward_demodulation, [], [f10430, f2655])).
fof(f2655, plain, ((e12 = op1(e11, e15)) | ~ spl42_255), inference(avatar_component_clause, [], [f2653])).
fof(f2653, plain, (spl42_255 <=> (e12 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_255])])).
fof(f16997, plain, (~ spl42_160 | ~ spl42_111), inference(avatar_split_clause, [], [f16069, f2056, f2259])).
fof(f2259, plain, (spl42_160 <=> (e15 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_160])])).
fof(f16069, plain, (~ (e15 = op1(e13, e15)) | ~ spl42_111), inference(forward_demodulation, [], [f623, f2058])).
fof(f623, plain, ~ (op1(e13, e15) = op1(e14, e15)), inference(cnf_transformation, [], [f7])).
fof(f16964, plain, (~ spl42_75 | ~ spl42_89), inference(avatar_split_clause, [], [f16963, f1965, f1907])).
fof(f1907, plain, (spl42_75 <=> (e14 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_75])])).
fof(f16963, plain, (~ (e14 = op1(e15, e13)) | ~ spl42_89), inference(forward_demodulation, [], [f762, f1967])).
fof(f762, plain, ~ (op1(e15, e11) = op1(e15, e13)), inference(cnf_transformation, [], [f7])).
fof(f16931, plain, (~ spl42_13 | ~ spl42_111), inference(avatar_split_clause, [], [f15129, f2056, f1650])).
fof(f1650, plain, (spl42_13 <=> (e15 = op1(e16, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_13])])).
fof(f15129, plain, (~ (e15 = op1(e16, e15)) | ~ spl42_111), inference(forward_demodulation, [], [f627, f2058])).
fof(f627, plain, ~ (op1(e14, e15) = op1(e16, e15)), inference(cnf_transformation, [], [f7])).
fof(f16860, plain, (~ spl42_56 | ~ spl42_105), inference(avatar_split_clause, [], [f16854, f2031, f1828])).
fof(f2031, plain, (spl42_105 <=> (e16 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_105])])).
fof(f16854, plain, (~ (e16 = op1(e15, e16)) | ~ spl42_105), inference(backward_demodulation, [], [f647, f2033])).
fof(f2033, plain, ((e16 = op1(e14, e16)) | ~ spl42_105), inference(avatar_component_clause, [], [f2031])).
fof(f647, plain, ~ (op1(e14, e16) = op1(e15, e16)), inference(cnf_transformation, [], [f7])).
fof(f16849, plain, (~ spl42_115 | ~ spl42_124 | ~ spl42_211), inference(avatar_contradiction_clause, [], [f16848])).
fof(f16848, plain, ($false | (~ spl42_115 | ~ spl42_124 | ~ spl42_211)), inference(subsumption_resolution, [], [f16847, f1093])).
fof(f1093, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f9])).
fof(f16847, plain, ((e10 = e13) | (~ spl42_115 | ~ spl42_124 | ~ spl42_211)), inference(forward_demodulation, [], [f16846, f2473])).
fof(f16846, plain, ((e13 = op1(e12, e14)) | (~ spl42_115 | ~ spl42_124)), inference(forward_demodulation, [], [f16840, f2075])).
fof(f16840, plain, ((e13 = op1(op1(e14, e14), e14)) | ~ spl42_124), inference(backward_demodulation, [], [f283, f2112])).
fof(f2112, plain, ((e14 = op1(e14, e13)) | ~ spl42_124), inference(avatar_component_clause, [], [f2110])).
fof(f2110, plain, (spl42_124 <=> (e14 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_124])])).
fof(f16838, plain, (~ spl42_127 | ~ spl42_130), inference(avatar_contradiction_clause, [], [f16837])).
fof(f16837, plain, ($false | (~ spl42_127 | ~ spl42_130)), inference(subsumption_resolution, [], [f16836, f1093])).
fof(f16836, plain, ((e10 = e13) | (~ spl42_127 | ~ spl42_130)), inference(forward_demodulation, [], [f2137, f2125])).
fof(f2137, plain, ((e13 = op1(e14, e12)) | ~ spl42_130), inference(avatar_component_clause, [], [f2135])).
fof(f2135, plain, (spl42_130 <=> (e13 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_130])])).
fof(f16835, plain, (~ spl42_127 | ~ spl42_131), inference(avatar_contradiction_clause, [], [f16834])).
fof(f16834, plain, ($false | (~ spl42_127 | ~ spl42_131)), inference(subsumption_resolution, [], [f16833, f1094])).
fof(f16833, plain, ((e10 = e14) | (~ spl42_127 | ~ spl42_131)), inference(forward_demodulation, [], [f2141, f2125])).
fof(f2141, plain, ((e14 = op1(e14, e12)) | ~ spl42_131), inference(avatar_component_clause, [], [f2139])).
fof(f2139, plain, (spl42_131 <=> (e14 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_131])])).
fof(f16737, plain, (~ spl42_188 | ~ spl42_237), inference(avatar_split_clause, [], [f16732, f2578, f2375])).
fof(f2578, plain, (spl42_237 <=> (e15 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_237])])).
fof(f16732, plain, (~ (e15 = op1(e13, e11)) | ~ spl42_237), inference(backward_demodulation, [], [f535, f2580])).
fof(f2580, plain, ((e15 = op1(e12, e11)) | ~ spl42_237), inference(avatar_component_clause, [], [f2578])).
fof(f535, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f16727, plain, (~ spl42_127 | ~ spl42_243 | spl42_323), inference(avatar_contradiction_clause, [], [f16726])).
fof(f16726, plain, ($false | (~ spl42_127 | ~ spl42_243 | spl42_323)), inference(subsumption_resolution, [], [f16725, f2936])).
fof(f2936, plain, (~ (e10 = op1(e10, e12)) | spl42_323), inference(avatar_component_clause, [], [f2935])).
fof(f2935, plain, (spl42_323 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_323])])).
fof(f16725, plain, ((e10 = op1(e10, e12)) | (~ spl42_127 | ~ spl42_243)), inference(forward_demodulation, [], [f16719, f2125])).
fof(f16719, plain, ((e10 = op1(op1(e14, e12), e12)) | ~ spl42_243), inference(backward_demodulation, [], [f260, f2605])).
fof(f2605, plain, ((e14 = op1(e12, e10)) | ~ spl42_243), inference(avatar_component_clause, [], [f2603])).
fof(f2603, plain, (spl42_243 <=> (e14 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_243])])).
fof(f260, plain, (e10 = op1(op1(op1(e12, e10), e12), e12)), inference(cnf_transformation, [], [f3])).
fof(f16690, plain, (~ spl42_89 | spl42_136 | ~ spl42_279), inference(avatar_contradiction_clause, [], [f16689])).
fof(f16689, plain, ($false | (~ spl42_89 | spl42_136 | ~ spl42_279)), inference(subsumption_resolution, [], [f16688, f2161])).
fof(f2161, plain, (~ (e12 = op1(e14, e11)) | spl42_136), inference(avatar_component_clause, [], [f2160])).
fof(f2160, plain, (spl42_136 <=> (e12 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_136])])).
fof(f16688, plain, ((e12 = op1(e14, e11)) | (~ spl42_89 | ~ spl42_279)), inference(forward_demodulation, [], [f16682, f1967])).
fof(f16682, plain, ((e12 = op1(op1(e15, e11), e11)) | ~ spl42_279), inference(backward_demodulation, [], [f273, f2754])).
fof(f2754, plain, ((e15 = op1(e11, e12)) | ~ spl42_279), inference(avatar_component_clause, [], [f2752])).
fof(f273, plain, (e12 = op1(op1(op1(e11, e12), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f16538, plain, (~ spl42_424 | ~ spl42_438), inference(avatar_split_clause, [], [f16529, f3509, f3451])).
fof(f3509, plain, (spl42_438 <=> (e23 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_438])])).
fof(f16529, plain, (~ (e23 = op2(e25, e22)) | ~ spl42_438), inference(backward_demodulation, [], [f1050, f3511])).
fof(f3511, plain, ((e23 = op2(e25, e20)) | ~ spl42_438), inference(avatar_component_clause, [], [f3509])).
fof(f1050, plain, ~ (op2(e25, e20) = op2(e25, e22)), inference(cnf_transformation, [], [f8])).
fof(f16516, plain, (~ spl42_432 | spl42_450 | ~ spl42_453), inference(avatar_contradiction_clause, [], [f16515])).
fof(f16515, plain, ($false | (~ spl42_432 | spl42_450 | ~ spl42_453)), inference(subsumption_resolution, [], [f16508, f3560])).
fof(f3560, plain, (~ (e21 = op2(e24, e25)) | spl42_450), inference(avatar_component_clause, [], [f3559])).
fof(f3559, plain, (spl42_450 <=> (e21 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_450])])).
fof(f16508, plain, ((e21 = op2(e24, e25)) | (~ spl42_432 | ~ spl42_453)), inference(backward_demodulation, [], [f12593, f3573])).
fof(f3573, plain, ((e24 = op2(e24, e25)) | ~ spl42_453), inference(avatar_component_clause, [], [f3571])).
fof(f3571, plain, (spl42_453 <=> (e24 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_453])])).
fof(f16488, plain, (~ spl42_470 | ~ spl42_475), inference(avatar_contradiction_clause, [], [f16487])).
fof(f16487, plain, ($false | (~ spl42_470 | ~ spl42_475)), inference(subsumption_resolution, [], [f16483, f1116])).
fof(f16483, plain, ((e20 = e25) | (~ spl42_470 | ~ spl42_475)), inference(backward_demodulation, [], [f3664, f3644])).
fof(f3664, plain, ((e25 = op2(e24, e22)) | ~ spl42_475), inference(avatar_component_clause, [], [f3662])).
fof(f3662, plain, (spl42_475 <=> (e25 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_475])])).
fof(f16454, plain, (spl42_415 | ~ spl42_517 | ~ spl42_529), inference(avatar_contradiction_clause, [], [f16453])).
fof(f16453, plain, ($false | (spl42_415 | ~ spl42_517 | ~ spl42_529)), inference(subsumption_resolution, [], [f16452, f3415])).
fof(f3415, plain, (~ (e21 = op2(e25, e23)) | spl42_415), inference(avatar_component_clause, [], [f3414])).
fof(f3414, plain, (spl42_415 <=> (e21 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_415])])).
fof(f16452, plain, ((e21 = op2(e25, e23)) | (~ spl42_517 | ~ spl42_529)), inference(backward_demodulation, [], [f16437, f3838])).
fof(f3838, plain, ((e25 = op2(e23, e23)) | ~ spl42_517), inference(avatar_component_clause, [], [f3836])).
fof(f3836, plain, (spl42_517 <=> (e25 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_517])])).
fof(f16355, plain, (spl42_350 | ~ spl42_399 | ~ spl42_692), inference(avatar_contradiction_clause, [], [f16354])).
fof(f16354, plain, ($false | (spl42_350 | ~ spl42_399 | ~ spl42_692)), inference(subsumption_resolution, [], [f16353, f3145])).
fof(f3145, plain, (~ (e26 = op2(e26, e26)) | spl42_350), inference(avatar_component_clause, [], [f3144])).
fof(f3144, plain, (spl42_350 <=> (e26 = op2(e26, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_350])])).
fof(f16353, plain, ((e26 = op2(e26, e26)) | (~ spl42_399 | ~ spl42_692)), inference(forward_demodulation, [], [f16342, f3349])).
fof(f16342, plain, ((e26 = op2(op2(e25, e26), e26)) | ~ spl42_692), inference(backward_demodulation, [], [f4826, f4870])).
fof(f4870, plain, ((e25 = h7(e11)) | ~ spl42_692), inference(avatar_component_clause, [], [f4869])).
fof(f4869, plain, (spl42_692 <=> (e25 = h7(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_692])])).
fof(f4826, plain, (e26 = op2(op2(h7(e11), e26), e26)), inference(backward_demodulation, [], [f502, f1252])).
fof(f502, plain, (e26 = op2(op2(op2(e26, e26), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f16307, plain, (~ spl42_531 | ~ spl42_794), inference(avatar_split_clause, [], [f16299, f5414, f3894])).
fof(f5414, plain, (spl42_794 <=> (e25 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_794])])).
fof(f16299, plain, (~ (e25 = op2(e23, e21)) | ~ spl42_794), inference(backward_demodulation, [], [f4751, f5415])).
fof(f5415, plain, ((e25 = h4(e11)) | ~ spl42_794), inference(avatar_component_clause, [], [f5414])).
fof(f4751, plain, ~ (op2(e23, e21) = h4(e11)), inference(backward_demodulation, [], [f1014, f1231])).
fof(f1014, plain, ~ (op2(e23, e21) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f16239, plain, (spl42_1336 | ~ spl42_686), inference(avatar_split_clause, [], [f15491, f4536, f8710])).
fof(f4536, plain, (spl42_686 <=> (op2(e20, e20) = e26)), introduced(avatar_definition, [new_symbols(naming, [spl42_686])])).
fof(f15491, plain, ((e26 = h1(e11)) | ~ spl42_686), inference(backward_demodulation, [], [f1210, f4538])).
fof(f4538, plain, ((op2(e20, e20) = e26) | ~ spl42_686), inference(avatar_component_clause, [], [f4536])).
fof(f16230, plain, (~ spl42_591 | ~ spl42_619), inference(avatar_split_clause, [], [f16229, f4259, f4143])).
fof(f4259, plain, (spl42_619 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_619])])).
fof(f16229, plain, (~ (e22 = op2(e21, e26)) | ~ spl42_619), inference(forward_demodulation, [], [f979, f4261])).
fof(f4261, plain, ((e22 = op2(e21, e22)) | ~ spl42_619), inference(avatar_component_clause, [], [f4259])).
fof(f979, plain, ~ (op2(e21, e22) = op2(e21, e26)), inference(cnf_transformation, [], [f8])).
fof(f16216, plain, (~ spl42_587 | ~ spl42_566), inference(avatar_split_clause, [], [f16215, f4039, f4126])).
fof(f4126, plain, (spl42_587 <=> (e25 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_587])])).
fof(f16215, plain, (~ (e25 = op2(e22, e20)) | ~ spl42_566), inference(forward_demodulation, [], [f988, f4041])).
fof(f4041, plain, ((e25 = op2(e22, e23)) | ~ spl42_566), inference(avatar_component_clause, [], [f4039])).
fof(f988, plain, ~ (op2(e22, e20) = op2(e22, e23)), inference(cnf_transformation, [], [f8])).
fof(f16181, plain, (~ spl42_430 | ~ spl42_432), inference(avatar_contradiction_clause, [], [f16180])).
fof(f16180, plain, ($false | (~ spl42_430 | ~ spl42_432)), inference(subsumption_resolution, [], [f16179, f1124])).
fof(f1124, plain, ~ (e22 = e24), inference(cnf_transformation, [], [f10])).
fof(f16179, plain, ((e22 = e24) | (~ spl42_430 | ~ spl42_432)), inference(forward_demodulation, [], [f3486, f3478])).
fof(f3478, plain, ((e22 = op2(e25, e21)) | ~ spl42_430), inference(avatar_component_clause, [], [f3476])).
fof(f3476, plain, (spl42_430 <=> (e22 = op2(e25, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_430])])).
fof(f16158, plain, (~ spl42_364 | ~ spl42_1016), inference(avatar_split_clause, [], [f13505, f6684, f3202])).
fof(f3202, plain, (spl42_364 <=> (e26 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_364])])).
fof(f13505, plain, (~ (e26 = op2(e26, e24)) | ~ spl42_1016), inference(forward_demodulation, [], [f4791, f6685])).
fof(f4791, plain, ~ (op2(e26, e24) = h6(e16)), inference(backward_demodulation, [], [f898, f4785])).
fof(f4785, plain, (op2(e23, e24) = h6(e16)), inference(forward_demodulation, [], [f4784, f4778])).
fof(f4784, plain, (h6(e16) = op2(op2(h5(e11), e25), e24)), inference(forward_demodulation, [], [f4783, f1238])).
fof(f4783, plain, (h6(e16) = op2(op2(op2(e24, e24), e25), e24)), inference(forward_demodulation, [], [f1249, f1206])).
fof(f1249, plain, (op2(op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), e25), op2(e25, op2(e25, e25))) = h6(e16)), inference(cnf_transformation, [], [f21])).
fof(f898, plain, ~ (op2(e23, e24) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f16129, plain, (spl42_184 | ~ spl42_268), inference(avatar_split_clause, [], [f16128, f2707, f2359])).
fof(f2359, plain, (spl42_184 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_184])])).
fof(f2707, plain, (spl42_268 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_268])])).
fof(f16128, plain, ((e11 = op1(e13, e11)) | ~ spl42_268), inference(forward_demodulation, [], [f266, f16119])).
fof(f16119, plain, ((e13 = op1(op1(e11, e11), e11)) | ~ spl42_268), inference(forward_demodulation, [], [f280, f2709])).
fof(f2709, plain, ((e11 = op1(e11, e13)) | ~ spl42_268), inference(avatar_component_clause, [], [f2707])).
fof(f280, plain, (e13 = op1(op1(op1(e11, e13), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f16092, plain, (~ spl42_80 | spl42_227 | ~ spl42_230), inference(avatar_contradiction_clause, [], [f16091])).
fof(f16091, plain, ($false | (~ spl42_80 | spl42_227 | ~ spl42_230)), inference(subsumption_resolution, [], [f16090, f2538])).
fof(f2538, plain, (~ (e12 = op1(e12, e12)) | spl42_227), inference(avatar_component_clause, [], [f2537])).
fof(f2537, plain, (spl42_227 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_227])])).
fof(f16090, plain, ((e12 = op1(e12, e12)) | (~ spl42_80 | ~ spl42_230)), inference(forward_demodulation, [], [f15807, f1930])).
fof(f1930, plain, ((e12 = op1(e15, e12)) | ~ spl42_80), inference(avatar_component_clause, [], [f1928])).
fof(f1928, plain, (spl42_80 <=> (e12 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_80])])).
fof(f15807, plain, ((e12 = op1(op1(e15, e12), e12)) | ~ spl42_230), inference(backward_demodulation, [], [f274, f2551])).
fof(f16074, plain, (spl42_188 | ~ spl42_137 | ~ spl42_257), inference(avatar_split_clause, [], [f15872, f2661, f2164, f2375])).
fof(f15872, plain, ((e15 = op1(e13, e11)) | (~ spl42_137 | ~ spl42_257)), inference(backward_demodulation, [], [f14969, f2166])).
fof(f14969, plain, ((e15 = op1(op1(e14, e11), e11)) | ~ spl42_257), inference(backward_demodulation, [], [f294, f2663])).
fof(f294, plain, (e15 = op1(op1(op1(e11, e15), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f16068, plain, (~ spl42_159 | ~ spl42_257), inference(avatar_split_clause, [], [f16067, f2661, f2255])).
fof(f2255, plain, (spl42_159 <=> (e14 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_159])])).
fof(f16067, plain, (~ (e14 = op1(e13, e15)) | ~ spl42_257), inference(forward_demodulation, [], [f615, f2663])).
fof(f615, plain, ~ (op1(e11, e15) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f16032, plain, (spl42_11 | ~ spl42_71 | ~ spl42_308), inference(avatar_split_clause, [], [f16031, f2872, f1891, f1642])).
fof(f1642, plain, (spl42_11 <=> (e13 = op1(e16, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_11])])).
fof(f16031, plain, ((e13 = op1(e16, e15)) | (~ spl42_71 | ~ spl42_308)), inference(forward_demodulation, [], [f16030, f2874])).
fof(f16030, plain, ((e13 = op1(op1(e10, e15), e15)) | ~ spl42_71), inference(forward_demodulation, [], [f284, f1893])).
fof(f284, plain, (e13 = op1(op1(op1(e15, e13), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f16016, plain, (~ spl42_30 | ~ spl42_16), inference(avatar_split_clause, [], [f15142, f1663, f1721])).
fof(f1721, plain, (spl42_30 <=> (e11 = op1(e16, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_30])])).
fof(f15142, plain, (~ (e11 = op1(e16, e12)) | ~ spl42_16), inference(forward_demodulation, [], [f788, f1665])).
fof(f788, plain, ~ (op1(e16, e12) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f15885, plain, (~ spl42_127 | ~ spl42_128), inference(avatar_contradiction_clause, [], [f15884])).
fof(f15884, plain, ($false | (~ spl42_127 | ~ spl42_128)), inference(subsumption_resolution, [], [f15883, f1091])).
fof(f15883, plain, ((e10 = e11) | (~ spl42_127 | ~ spl42_128)), inference(forward_demodulation, [], [f2129, f2125])).
fof(f2129, plain, ((e11 = op1(e14, e12)) | ~ spl42_128), inference(avatar_component_clause, [], [f2127])).
fof(f2127, plain, (spl42_128 <=> (e11 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_128])])).
fof(f15854, plain, (~ spl42_165 | ~ spl42_168), inference(avatar_contradiction_clause, [], [f15853])).
fof(f15853, plain, ($false | (~ spl42_165 | ~ spl42_168)), inference(subsumption_resolution, [], [f15851, f1108])).
fof(f15851, plain, ((e13 = e16) | (~ spl42_165 | ~ spl42_168)), inference(backward_demodulation, [], [f2294, f2282])).
fof(f2282, plain, ((e13 = op1(e13, e14)) | ~ spl42_165), inference(avatar_component_clause, [], [f2280])).
fof(f2280, plain, (spl42_165 <=> (e13 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_165])])).
fof(f15715, plain, (spl42_348 | ~ spl42_363 | ~ spl42_399), inference(avatar_contradiction_clause, [], [f15714])).
fof(f15714, plain, ($false | (spl42_348 | ~ spl42_363 | ~ spl42_399)), inference(subsumption_resolution, [], [f15709, f3137])).
fof(f3137, plain, (~ (e24 = op2(e26, e26)) | spl42_348), inference(avatar_component_clause, [], [f3136])).
fof(f15709, plain, ((e24 = op2(e26, e26)) | (~ spl42_363 | ~ spl42_399)), inference(backward_demodulation, [], [f15331, f3349])).
fof(f15331, plain, ((e24 = op2(op2(e25, e26), e26)) | ~ spl42_363), inference(forward_demodulation, [], [f488, f3200])).
fof(f3200, plain, ((e25 = op2(e26, e24)) | ~ spl42_363), inference(avatar_component_clause, [], [f3198])).
fof(f3198, plain, (spl42_363 <=> (e25 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_363])])).
fof(f488, plain, (e24 = op2(op2(op2(e26, e24), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f15704, plain, (~ spl42_416 | spl42_501 | ~ spl42_550), inference(avatar_contradiction_clause, [], [f15703])).
fof(f15703, plain, ($false | (~ spl42_416 | spl42_501 | ~ spl42_550)), inference(subsumption_resolution, [], [f15702, f3771])).
fof(f3771, plain, (~ (e23 = op2(e23, e25)) | spl42_501), inference(avatar_component_clause, [], [f3770])).
fof(f3770, plain, (spl42_501 <=> (e23 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_501])])).
fof(f15702, plain, ((e23 = op2(e23, e25)) | (~ spl42_416 | ~ spl42_550)), inference(forward_demodulation, [], [f15694, f3975])).
fof(f15694, plain, ((e23 = op2(op2(e22, e25), e25)) | ~ spl42_416), inference(backward_demodulation, [], [f480, f3420])).
fof(f3420, plain, ((e22 = op2(e25, e23)) | ~ spl42_416), inference(avatar_component_clause, [], [f3418])).
fof(f3418, plain, (spl42_416 <=> (e22 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_416])])).
fof(f480, plain, (e23 = op2(op2(op2(e25, e23), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f15690, plain, (~ spl42_431 | ~ spl42_432), inference(avatar_contradiction_clause, [], [f15689])).
fof(f15689, plain, ($false | (~ spl42_431 | ~ spl42_432)), inference(subsumption_resolution, [], [f15688, f1127])).
fof(f15688, plain, ((e23 = e24) | (~ spl42_431 | ~ spl42_432)), inference(backward_demodulation, [], [f3486, f3482])).
fof(f3482, plain, ((e23 = op2(e25, e21)) | ~ spl42_431), inference(avatar_component_clause, [], [f3480])).
fof(f3480, plain, (spl42_431 <=> (e23 = op2(e25, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_431])])).
fof(f15682, plain, (~ spl42_437 | spl42_498 | ~ spl42_550), inference(avatar_contradiction_clause, [], [f15681])).
fof(f15681, plain, ($false | (~ spl42_437 | spl42_498 | ~ spl42_550)), inference(subsumption_resolution, [], [f15680, f3759])).
fof(f3759, plain, (~ (e20 = op2(e23, e25)) | spl42_498), inference(avatar_component_clause, [], [f3758])).
fof(f15680, plain, ((e20 = op2(e23, e25)) | (~ spl42_437 | ~ spl42_550)), inference(forward_demodulation, [], [f15671, f3975])).
fof(f15671, plain, ((e20 = op2(op2(e22, e25), e25)) | ~ spl42_437), inference(backward_demodulation, [], [f459, f3507])).
fof(f459, plain, (e20 = op2(op2(op2(e25, e20), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f15647, plain, (~ spl42_458 | ~ spl42_459), inference(avatar_contradiction_clause, [], [f15646])).
fof(f15646, plain, ($false | (~ spl42_458 | ~ spl42_459)), inference(subsumption_resolution, [], [f15645, f1123])).
fof(f15645, plain, ((e22 = e23) | (~ spl42_458 | ~ spl42_459)), inference(forward_demodulation, [], [f3598, f3594])).
fof(f3594, plain, ((e22 = op2(e24, e24)) | ~ spl42_458), inference(avatar_component_clause, [], [f3592])).
fof(f3592, plain, (spl42_458 <=> (e22 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_458])])).
fof(f3598, plain, ((e23 = op2(e24, e24)) | ~ spl42_459), inference(avatar_component_clause, [], [f3596])).
fof(f3596, plain, (spl42_459 <=> (e23 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_459])])).
fof(f15642, plain, (~ spl42_458 | ~ spl42_463 | ~ spl42_656), inference(avatar_contradiction_clause, [], [f15641])).
fof(f15641, plain, ($false | (~ spl42_458 | ~ spl42_463 | ~ spl42_656)), inference(subsumption_resolution, [], [f15640, f1123])).
fof(f15640, plain, ((e22 = e23) | (~ spl42_458 | ~ spl42_463 | ~ spl42_656)), inference(forward_demodulation, [], [f15639, f3594])).
fof(f15639, plain, ((e23 = op2(e24, e24)) | (~ spl42_463 | ~ spl42_656)), inference(forward_demodulation, [], [f15627, f4414])).
fof(f15627, plain, ((e23 = op2(op2(e20, e24), e24)) | ~ spl42_463), inference(backward_demodulation, [], [f479, f3615])).
fof(f3615, plain, ((e20 = op2(e24, e23)) | ~ spl42_463), inference(avatar_component_clause, [], [f3613])).
fof(f479, plain, (e23 = op2(op2(op2(e24, e23), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f15616, plain, (~ spl42_508 | ~ spl42_511), inference(avatar_contradiction_clause, [], [f15615])).
fof(f15615, plain, ($false | (~ spl42_508 | ~ spl42_511)), inference(subsumption_resolution, [], [f15614, f1129])).
fof(f15614, plain, ((e23 = e26) | (~ spl42_508 | ~ spl42_511)), inference(backward_demodulation, [], [f3813, f3801])).
fof(f3801, plain, ((e23 = op2(e23, e24)) | ~ spl42_508), inference(avatar_component_clause, [], [f3799])).
fof(f3799, plain, (spl42_508 <=> (e23 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_508])])).
fof(f15609, plain, (spl42_475 | ~ spl42_523 | ~ spl42_550), inference(avatar_split_clause, [], [f15605, f3973, f3861, f3662])).
fof(f3861, plain, (spl42_523 <=> (e24 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_523])])).
fof(f15605, plain, ((e25 = op2(e24, e22)) | (~ spl42_523 | ~ spl42_550)), inference(backward_demodulation, [], [f13666, f3863])).
fof(f3863, plain, ((e24 = op2(e23, e22)) | ~ spl42_523), inference(avatar_component_clause, [], [f3861])).
fof(f15598, plain, (spl42_431 | ~ spl42_531 | ~ spl42_613), inference(avatar_split_clause, [], [f15597, f4234, f3894, f3480])).
fof(f15597, plain, ((e23 = op2(e25, e21)) | (~ spl42_531 | ~ spl42_613)), inference(backward_demodulation, [], [f15544, f3896])).
fof(f15544, plain, ((e23 = op2(op2(e23, e21), e21)) | ~ spl42_613), inference(backward_demodulation, [], [f476, f4236])).
fof(f4236, plain, ((e23 = op2(e21, e23)) | ~ spl42_613), inference(avatar_component_clause, [], [f4234])).
fof(f476, plain, (e23 = op2(op2(op2(e21, e23), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f15581, plain, (~ spl42_531 | ~ spl42_580), inference(avatar_split_clause, [], [f15573, f4097, f3894])).
fof(f4097, plain, (spl42_580 <=> (e25 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_580])])).
fof(f15573, plain, (~ (e25 = op2(e23, e21)) | ~ spl42_580), inference(backward_demodulation, [], [f829, f4099])).
fof(f4099, plain, ((e25 = op2(e22, e21)) | ~ spl42_580), inference(avatar_component_clause, [], [f4097])).
fof(f829, plain, ~ (op2(e22, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f8])).
fof(f15552, plain, (~ spl42_466 | ~ spl42_613), inference(avatar_split_clause, [], [f15545, f4234, f3625])).
fof(f3625, plain, (spl42_466 <=> (e23 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_466])])).
fof(f15545, plain, (~ (e23 = op2(e24, e23)) | ~ spl42_613), inference(backward_demodulation, [], [f868, f4236])).
fof(f15510, plain, (spl42_523 | ~ spl42_554 | ~ spl42_669), inference(avatar_split_clause, [], [f15503, f4466, f3990, f3861])).
fof(f3990, plain, (spl42_554 <=> (e20 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_554])])).
fof(f15503, plain, ((e24 = op2(e23, e22)) | (~ spl42_554 | ~ spl42_669)), inference(backward_demodulation, [], [f12834, f4468])).
fof(f4468, plain, ((e23 = op2(e20, e22)) | ~ spl42_669), inference(avatar_component_clause, [], [f4466])).
fof(f12834, plain, ((e24 = op2(op2(e20, e22), e22)) | ~ spl42_554), inference(backward_demodulation, [], [f484, f3992])).
fof(f3992, plain, ((e20 = op2(e22, e24)) | ~ spl42_554), inference(avatar_component_clause, [], [f3990])).
fof(f484, plain, (e24 = op2(op2(op2(e22, e24), e22), e22)), inference(cnf_transformation, [], [f6])).
fof(f15443, plain, (~ spl42_631 | ~ spl42_637), inference(avatar_contradiction_clause, [], [f15442])).
fof(f15442, plain, ($false | (~ spl42_631 | ~ spl42_637)), inference(subsumption_resolution, [], [f15441, f1117])).
fof(f1117, plain, ~ (e20 = e26), inference(cnf_transformation, [], [f10])).
fof(f15441, plain, ((e20 = e26) | (~ spl42_631 | ~ spl42_637)), inference(forward_demodulation, [], [f4335, f4311])).
fof(f4335, plain, ((e26 = op2(e21, e20)) | ~ spl42_637), inference(avatar_component_clause, [], [f4333])).
fof(f4333, plain, (spl42_637 <=> (e26 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_637])])).
fof(f15362, plain, (spl42_467 | ~ spl42_369 | ~ spl42_1016), inference(avatar_split_clause, [], [f14895, f6684, f3223, f3629])).
fof(f3629, plain, (spl42_467 <=> (e24 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_467])])).
fof(f3223, plain, (spl42_369 <=> (e24 = op2(e26, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_369])])).
fof(f14895, plain, ((e24 = op2(e24, e23)) | (~ spl42_369 | ~ spl42_1016)), inference(backward_demodulation, [], [f11251, f3225])).
fof(f3225, plain, ((e24 = op2(e26, e23)) | ~ spl42_369), inference(avatar_component_clause, [], [f3223])).
fof(f11251, plain, ((e24 = op2(op2(e26, e23), e23)) | ~ spl42_1016), inference(backward_demodulation, [], [f4786, f6685])).
fof(f4786, plain, (e24 = op2(op2(h6(e16), e23), e23)), inference(backward_demodulation, [], [f485, f4785])).
fof(f485, plain, (e24 = op2(op2(op2(e23, e24), e23), e23)), inference(cnf_transformation, [], [f6])).
fof(f15358, plain, (~ spl42_437 | ~ spl42_584), inference(avatar_split_clause, [], [f15357, f4114, f3505])).
fof(f4114, plain, (spl42_584 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_584])])).
fof(f15357, plain, (~ (e22 = op2(e25, e20)) | ~ spl42_584), inference(forward_demodulation, [], [f810, f4116])).
fof(f4116, plain, ((e22 = op2(e22, e20)) | ~ spl42_584), inference(avatar_component_clause, [], [f4114])).
fof(f810, plain, ~ (op2(e22, e20) = op2(e25, e20)), inference(cnf_transformation, [], [f8])).
fof(f15317, plain, (~ spl42_286 | ~ spl42_188), inference(avatar_split_clause, [], [f15316, f2375, f2781])).
fof(f2781, plain, (spl42_286 <=> (e15 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_286])])).
fof(f15316, plain, (~ (e15 = op1(e11, e11)) | ~ spl42_188), inference(forward_demodulation, [], [f531, f2377])).
fof(f531, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f15261, plain, (~ spl42_236 | ~ spl42_89), inference(avatar_split_clause, [], [f15260, f1965, f2574])).
fof(f2574, plain, (spl42_236 <=> (e14 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_236])])).
fof(f15260, plain, (~ (e14 = op1(e12, e11)) | ~ spl42_89), inference(forward_demodulation, [], [f537, f1967])).
fof(f537, plain, ~ (op1(e12, e11) = op1(e15, e11)), inference(cnf_transformation, [], [f7])).
fof(f15256, plain, (~ spl42_235 | ~ spl42_207), inference(avatar_split_clause, [], [f15255, f2454, f2570])).
fof(f2570, plain, (spl42_235 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_235])])).
fof(f15255, plain, (~ (e13 = op1(e12, e11)) | ~ spl42_207), inference(forward_demodulation, [], [f701, f2456])).
fof(f701, plain, ~ (op1(e12, e11) = op1(e12, e15)), inference(cnf_transformation, [], [f7])).
fof(f15207, plain, (~ spl42_127 | ~ spl42_180 | ~ spl42_207), inference(avatar_contradiction_clause, [], [f15206])).
fof(f15206, plain, ($false | (~ spl42_127 | ~ spl42_180 | ~ spl42_207)), inference(subsumption_resolution, [], [f15205, f1095])).
fof(f15205, plain, ((e10 = e15) | (~ spl42_127 | ~ spl42_180 | ~ spl42_207)), inference(forward_demodulation, [], [f15021, f2125])).
fof(f15021, plain, ((e15 = op1(e14, e12)) | (~ spl42_180 | ~ spl42_207)), inference(backward_demodulation, [], [f13407, f2344])).
fof(f2344, plain, ((e14 = op1(e13, e12)) | ~ spl42_180), inference(avatar_component_clause, [], [f2342])).
fof(f2342, plain, (spl42_180 <=> (e14 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_180])])).
fof(f15204, plain, (~ spl42_123 | ~ spl42_74), inference(avatar_split_clause, [], [f15203, f1903, f2106])).
fof(f1903, plain, (spl42_74 <=> (e13 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_74])])).
fof(f15203, plain, (~ (e13 = op1(e14, e13)) | ~ spl42_74), inference(forward_demodulation, [], [f584, f1905])).
fof(f1905, plain, ((e13 = op1(e15, e13)) | ~ spl42_74), inference(avatar_component_clause, [], [f1903])).
fof(f584, plain, ~ (op1(e14, e13) = op1(e15, e13)), inference(cnf_transformation, [], [f7])).
fof(f15192, plain, (~ spl42_104 | ~ spl42_111), inference(avatar_split_clause, [], [f15057, f2056, f2027])).
fof(f2027, plain, (spl42_104 <=> (e15 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_104])])).
fof(f15057, plain, (~ (e15 = op1(e14, e16)) | ~ spl42_111), inference(backward_demodulation, [], [f754, f2058])).
fof(f754, plain, ~ (op1(e14, e15) = op1(e14, e16)), inference(cnf_transformation, [], [f7])).
fof(f15173, plain, (spl42_305 | ~ spl42_74 | ~ spl42_155), inference(avatar_split_clause, [], [f15172, f2239, f1903, f2860])).
fof(f2860, plain, (spl42_305 <=> (e13 = op1(e10, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_305])])).
fof(f15172, plain, ((e13 = op1(e10, e15)) | (~ spl42_74 | ~ spl42_155)), inference(forward_demodulation, [], [f15171, f2241])).
fof(f15171, plain, ((e13 = op1(op1(e13, e15), e15)) | ~ spl42_74), inference(forward_demodulation, [], [f284, f1905])).
fof(f15126, plain, (~ spl42_11 | ~ spl42_207), inference(avatar_split_clause, [], [f15125, f2454, f1642])).
fof(f15125, plain, (~ (e13 = op1(e16, e15)) | ~ spl42_207), inference(forward_demodulation, [], [f622, f2456])).
fof(f622, plain, ~ (op1(e12, e15) = op1(e16, e15)), inference(cnf_transformation, [], [f7])).
fof(f15084, plain, (~ spl42_58 | ~ spl42_59), inference(avatar_contradiction_clause, [], [f15083])).
fof(f15083, plain, ($false | (~ spl42_58 | ~ spl42_59)), inference(subsumption_resolution, [], [f15082, f1097])).
fof(f1097, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f9])).
fof(f15082, plain, ((e11 = e12) | (~ spl42_58 | ~ spl42_59)), inference(forward_demodulation, [], [f1843, f1839])).
fof(f1843, plain, ((e12 = op1(e15, e15)) | ~ spl42_59), inference(avatar_component_clause, [], [f1841])).
fof(f1841, plain, (spl42_59 <=> (e12 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_59])])).
fof(f15081, plain, (~ spl42_20 | ~ spl42_69), inference(avatar_split_clause, [], [f15077, f1882, f1679])).
fof(f1679, plain, (spl42_20 <=> (e15 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_20])])).
fof(f15077, plain, (~ (e15 = op1(e16, e14)) | ~ spl42_69), inference(backward_demodulation, [], [f607, f1884])).
fof(f607, plain, ~ (op1(e15, e14) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f15067, plain, (spl42_21 | ~ spl42_102 | ~ spl42_168), inference(avatar_contradiction_clause, [], [f15066])).
fof(f15066, plain, ($false | (spl42_21 | ~ spl42_102 | ~ spl42_168)), inference(subsumption_resolution, [], [f15065, f1684])).
fof(f1684, plain, (~ (e16 = op1(e16, e14)) | spl42_21), inference(avatar_component_clause, [], [f1683])).
fof(f1683, plain, (spl42_21 <=> (e16 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_21])])).
fof(f15065, plain, ((e16 = op1(e16, e14)) | (~ spl42_102 | ~ spl42_168)), inference(forward_demodulation, [], [f15060, f2294])).
fof(f15060, plain, ((e16 = op1(op1(e13, e14), e14)) | ~ spl42_102), inference(backward_demodulation, [], [f304, f2021])).
fof(f2021, plain, ((e13 = op1(e14, e16)) | ~ spl42_102), inference(avatar_component_clause, [], [f2019])).
fof(f2019, plain, (spl42_102 <=> (e13 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_102])])).
fof(f304, plain, (e16 = op1(op1(op1(e14, e16), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f15014, plain, (~ spl42_181 | ~ spl42_188), inference(avatar_split_clause, [], [f15006, f2375, f2346])).
fof(f2346, plain, (spl42_181 <=> (e15 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_181])])).
fof(f15006, plain, (~ (e15 = op1(e13, e12)) | ~ spl42_188), inference(backward_demodulation, [], [f719, f2377])).
fof(f719, plain, ~ (op1(e13, e11) = op1(e13, e12)), inference(cnf_transformation, [], [f7])).
fof(f14982, plain, (~ spl42_248 | ~ spl42_251), inference(avatar_contradiction_clause, [], [f14981])).
fof(f14981, plain, ($false | (~ spl42_248 | ~ spl42_251)), inference(subsumption_resolution, [], [f14980, f1104])).
fof(f14980, plain, ((e12 = e15) | (~ spl42_248 | ~ spl42_251)), inference(forward_demodulation, [], [f2638, f2626])).
fof(f2638, plain, ((e15 = op1(e11, e16)) | ~ spl42_251), inference(avatar_component_clause, [], [f2636])).
fof(f2636, plain, (spl42_251 <=> (e15 = op1(e11, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_251])])).
fof(f14922, plain, (~ spl42_349 | spl42_692), inference(avatar_contradiction_clause, [], [f14921])).
fof(f14921, plain, ($false | (~ spl42_349 | spl42_692)), inference(subsumption_resolution, [], [f14920, f4871])).
fof(f4871, plain, (~ (e25 = h7(e11)) | spl42_692), inference(avatar_component_clause, [], [f4869])).
fof(f14920, plain, ((e25 = h7(e11)) | ~ spl42_349), inference(backward_demodulation, [], [f1252, f3142])).
fof(f3142, plain, ((e25 = op2(e26, e26)) | ~ spl42_349), inference(avatar_component_clause, [], [f3140])).
fof(f3140, plain, (spl42_349 <=> (e25 = op2(e26, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_349])])).
fof(f14879, plain, (~ spl42_378 | ~ spl42_544), inference(avatar_contradiction_clause, [], [f14878])).
fof(f14878, plain, ($false | (~ spl42_378 | ~ spl42_544)), inference(subsumption_resolution, [], [f14877, f1131])).
fof(f1131, plain, ~ (e24 = e26), inference(cnf_transformation, [], [f10])).
fof(f14877, plain, ((e24 = e26) | (~ spl42_378 | ~ spl42_544)), inference(forward_demodulation, [], [f14876, f3950])).
fof(f14876, plain, ((e26 = op2(e22, e26)) | ~ spl42_378), inference(backward_demodulation, [], [f4826, f14875])).
fof(f14875, plain, ((e22 = op2(h7(e11), e26)) | ~ spl42_378), inference(forward_demodulation, [], [f14868, f1252])).
fof(f14868, plain, ((e22 = op2(op2(e26, e26), e26)) | ~ spl42_378), inference(backward_demodulation, [], [f474, f3262])).
fof(f3262, plain, ((e26 = op2(e26, e22)) | ~ spl42_378), inference(avatar_component_clause, [], [f3260])).
fof(f3260, plain, (spl42_378 <=> (e26 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_378])])).
fof(f474, plain, (e22 = op2(op2(op2(e26, e22), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f14827, plain, (~ spl42_356 | ~ spl42_454), inference(avatar_split_clause, [], [f14819, f3575, f3169])).
fof(f3169, plain, (spl42_356 <=> (e25 = op2(e26, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_356])])).
fof(f14819, plain, (~ (e25 = op2(e26, e25)) | ~ spl42_454), inference(backward_demodulation, [], [f921, f3577])).
fof(f921, plain, ~ (op2(e24, e25) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f14826, plain, (~ spl42_410 | ~ spl42_454 | spl42_510), inference(avatar_contradiction_clause, [], [f14825])).
fof(f14825, plain, ($false | (~ spl42_410 | ~ spl42_454 | spl42_510)), inference(subsumption_resolution, [], [f14824, f3808])).
fof(f3808, plain, (~ (e25 = op2(e23, e24)) | spl42_510), inference(avatar_component_clause, [], [f3807])).
fof(f3807, plain, (spl42_510 <=> (e25 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_510])])).
fof(f14824, plain, ((e25 = op2(e23, e24)) | (~ spl42_410 | ~ spl42_454)), inference(forward_demodulation, [], [f14818, f3395])).
fof(f3395, plain, ((e23 = op2(e25, e24)) | ~ spl42_410), inference(avatar_component_clause, [], [f3393])).
fof(f3393, plain, (spl42_410 <=> (e23 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_410])])).
fof(f14818, plain, ((e25 = op2(op2(e25, e24), e24)) | ~ spl42_454), inference(backward_demodulation, [], [f493, f3577])).
fof(f493, plain, (e25 = op2(op2(op2(e24, e25), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f14816, plain, (~ spl42_359 | ~ spl42_466 | ~ spl42_511), inference(avatar_contradiction_clause, [], [f14815])).
fof(f14815, plain, ($false | (~ spl42_359 | ~ spl42_466 | ~ spl42_511)), inference(subsumption_resolution, [], [f14814, f1119])).
fof(f1119, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f10])).
fof(f14814, plain, ((e21 = e23) | (~ spl42_359 | ~ spl42_466 | ~ spl42_511)), inference(forward_demodulation, [], [f14813, f3184])).
fof(f14813, plain, ((e23 = op2(e26, e24)) | (~ spl42_466 | ~ spl42_511)), inference(forward_demodulation, [], [f14804, f3813])).
fof(f14804, plain, ((e23 = op2(op2(e23, e24), e24)) | ~ spl42_466), inference(backward_demodulation, [], [f479, f3627])).
fof(f3627, plain, ((e23 = op2(e24, e23)) | ~ spl42_466), inference(avatar_component_clause, [], [f3625])).
fof(f14803, plain, (spl42_785 | ~ spl42_470 | ~ spl42_743), inference(avatar_split_clause, [], [f14799, f5142, f3642, f5368])).
fof(f14799, plain, ((e20 = h5(e14)) | (~ spl42_470 | ~ spl42_743)), inference(backward_demodulation, [], [f11293, f3644])).
fof(f11293, plain, ((op2(e24, e22) = h5(e14)) | ~ spl42_743), inference(backward_demodulation, [], [f4777, f5143])).
fof(f4777, plain, (h5(e14) = op2(e24, h5(e11))), inference(backward_demodulation, [], [f1241, f1238])).
fof(f1241, plain, (op2(e24, op2(e24, e24)) = h5(e14)), inference(cnf_transformation, [], [f20])).
fof(f14787, plain, (~ spl42_491 | ~ spl42_494), inference(avatar_contradiction_clause, [], [f14786])).
fof(f14786, plain, ($false | (~ spl42_491 | ~ spl42_494)), inference(subsumption_resolution, [], [f14785, f1114])).
fof(f14785, plain, ((e20 = e23) | (~ spl42_491 | ~ spl42_494)), inference(backward_demodulation, [], [f3743, f3731])).
fof(f3731, plain, ((e20 = op2(e23, e26)) | ~ spl42_491), inference(avatar_component_clause, [], [f3729])).
fof(f3729, plain, (spl42_491 <=> (e20 = op2(e23, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_491])])).
fof(f14772, plain, (~ spl42_517 | spl42_794), inference(avatar_contradiction_clause, [], [f14771])).
fof(f14771, plain, ($false | (~ spl42_517 | spl42_794)), inference(subsumption_resolution, [], [f14770, f5416])).
fof(f5416, plain, (~ (e25 = h4(e11)) | spl42_794), inference(avatar_component_clause, [], [f5414])).
fof(f14770, plain, ((e25 = h4(e11)) | ~ spl42_517), inference(backward_demodulation, [], [f1231, f3838])).
fof(f14708, plain, (spl42_836 | ~ spl42_573), inference(avatar_split_clause, [], [f14707, f4068, f5625])).
fof(f14707, plain, ((e25 = h3(e11)) | ~ spl42_573), inference(backward_demodulation, [], [f1224, f4070])).
fof(f14703, plain, (~ spl42_550 | ~ spl42_578), inference(avatar_contradiction_clause, [], [f14702])).
fof(f14702, plain, ($false | (~ spl42_550 | ~ spl42_578)), inference(subsumption_resolution, [], [f14701, f1121])).
fof(f14701, plain, ((e21 = e25) | (~ spl42_550 | ~ spl42_578)), inference(backward_demodulation, [], [f13666, f14693])).
fof(f14693, plain, ((e21 = op2(op2(e23, e22), e22)) | ~ spl42_578), inference(backward_demodulation, [], [f463, f4091])).
fof(f4091, plain, ((e23 = op2(e22, e21)) | ~ spl42_578), inference(avatar_component_clause, [], [f4089])).
fof(f4089, plain, (spl42_578 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_578])])).
fof(f463, plain, (e21 = op2(op2(op2(e22, e21), e22), e22)), inference(cnf_transformation, [], [f6])).
fof(f14671, plain, (spl42_480 | ~ spl42_432 | ~ spl42_615), inference(avatar_split_clause, [], [f14670, f4242, f3484, f3683])).
fof(f14670, plain, ((e23 = op2(e24, e21)) | (~ spl42_432 | ~ spl42_615)), inference(forward_demodulation, [], [f14664, f3486])).
fof(f14664, plain, ((e23 = op2(op2(e25, e21), e21)) | ~ spl42_615), inference(backward_demodulation, [], [f476, f4244])).
fof(f14590, plain, (~ spl42_613 | ~ spl42_620), inference(avatar_split_clause, [], [f14589, f4263, f4234])).
fof(f14589, plain, (~ (e23 = op2(e21, e23)) | ~ spl42_620), inference(forward_demodulation, [], [f976, f4265])).
fof(f4265, plain, ((e23 = op2(e21, e22)) | ~ spl42_620), inference(avatar_component_clause, [], [f4263])).
fof(f976, plain, ~ (op2(e21, e22) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f14588, plain, (~ spl42_632 | ~ spl42_673), inference(avatar_contradiction_clause, [], [f14587])).
fof(f14587, plain, ($false | (~ spl42_632 | ~ spl42_673)), inference(subsumption_resolution, [], [f14586, f1112])).
fof(f1112, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f10])).
fof(f14586, plain, ((e20 = e21) | (~ spl42_632 | ~ spl42_673)), inference(forward_demodulation, [], [f14585, f4485])).
fof(f14585, plain, ((e21 = op2(e20, e21)) | ~ spl42_632), inference(forward_demodulation, [], [f4699, f13709])).
fof(f13709, plain, ((e20 = op2(h2(e11), e21)) | ~ spl42_632), inference(forward_demodulation, [], [f13701, f1217])).
fof(f13701, plain, ((e20 = op2(op2(e21, e21), e21)) | ~ spl42_632), inference(backward_demodulation, [], [f455, f4315])).
fof(f4315, plain, ((e21 = op2(e21, e20)) | ~ spl42_632), inference(avatar_component_clause, [], [f4313])).
fof(f4313, plain, (spl42_632 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_632])])).
fof(f455, plain, (e20 = op2(op2(op2(e21, e20), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f4699, plain, (e21 = op2(op2(h2(e11), e21), e21)), inference(backward_demodulation, [], [f462, f1217])).
fof(f462, plain, (e21 = op2(op2(op2(e21, e21), e21), e21)), inference(cnf_transformation, [], [f6])).
fof(f14578, plain, (~ spl42_604 | ~ spl42_359), inference(avatar_split_clause, [], [f14577, f3182, f4197])).
fof(f4197, plain, (spl42_604 <=> (e21 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_604])])).
fof(f14577, plain, (~ (e21 = op2(e21, e24)) | ~ spl42_359), inference(forward_demodulation, [], [f891, f3184])).
fof(f891, plain, ~ (op2(e21, e24) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f14548, plain, (~ spl42_585 | ~ spl42_550), inference(avatar_split_clause, [], [f14547, f3973, f4118])).
fof(f14547, plain, (~ (e23 = op2(e22, e20)) | ~ spl42_550), inference(forward_demodulation, [], [f990, f3975])).
fof(f990, plain, ~ (op2(e22, e20) = op2(e22, e25)), inference(cnf_transformation, [], [f8])).
fof(f14512, plain, (spl42_504 | ~ spl42_395 | ~ spl42_550), inference(avatar_split_clause, [], [f13913, f3973, f3331, f3782])).
fof(f3782, plain, (spl42_504 <=> (e26 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_504])])).
fof(f3331, plain, (spl42_395 <=> (e22 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_395])])).
fof(f13913, plain, ((e26 = op2(e23, e25)) | (~ spl42_395 | ~ spl42_550)), inference(forward_demodulation, [], [f13906, f3975])).
fof(f13906, plain, ((e26 = op2(op2(e22, e25), e25)) | ~ spl42_395), inference(backward_demodulation, [], [f501, f3333])).
fof(f3333, plain, ((e22 = op2(e25, e26)) | ~ spl42_395), inference(avatar_component_clause, [], [f3331])).
fof(f14480, plain, (~ spl42_439 | ~ spl42_432), inference(avatar_split_clause, [], [f12604, f3484, f3513])).
fof(f3513, plain, (spl42_439 <=> (e24 = op2(e25, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_439])])).
fof(f12604, plain, (~ (e24 = op2(e25, e20)) | ~ spl42_432), inference(forward_demodulation, [], [f1049, f3486])).
fof(f1049, plain, ~ (op2(e25, e20) = op2(e25, e21)), inference(cnf_transformation, [], [f8])).
fof(f14429, plain, (~ spl42_352 | ~ spl42_748), inference(avatar_split_clause, [], [f12466, f5171, f3153])).
fof(f12466, plain, (~ (e21 = op2(e26, e25)) | ~ spl42_748), inference(forward_demodulation, [], [f4809, f5172])).
fof(f4809, plain, ~ (op2(e26, e25) = h6(e11)), inference(backward_demodulation, [], [f922, f1245])).
fof(f922, plain, ~ (op2(e25, e25) = op2(e26, e25)), inference(cnf_transformation, [], [f8])).
fof(f14418, plain, (spl42_319 | ~ spl42_193), inference(avatar_split_clause, [], [f13386, f2396, f2918])).
fof(f2396, plain, (spl42_193 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_193])])).
fof(f13386, plain, ((e13 = op1(e10, e13)) | ~ spl42_193), inference(forward_demodulation, [], [f282, f13097])).
fof(f13097, plain, ((e10 = op1(op1(e13, e13), e13)) | ~ spl42_193), inference(backward_demodulation, [], [f261, f2398])).
fof(f2398, plain, ((e13 = op1(e13, e10)) | ~ spl42_193), inference(avatar_component_clause, [], [f2396])).
fof(f261, plain, (e10 = op1(op1(op1(e13, e10), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f282, plain, (e13 = op1(op1(op1(e13, e13), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f14409, plain, (~ spl42_305 | ~ spl42_207), inference(avatar_split_clause, [], [f14408, f2454, f2860])).
fof(f14408, plain, (~ (e13 = op1(e10, e15)) | ~ spl42_207), inference(forward_demodulation, [], [f609, f2456])).
fof(f609, plain, ~ (op1(e10, e15) = op1(e12, e15)), inference(cnf_transformation, [], [f7])).
fof(f14396, plain, (spl42_8 | ~ spl42_92 | ~ spl42_308), inference(avatar_contradiction_clause, [], [f14395])).
fof(f14395, plain, ($false | (spl42_8 | ~ spl42_92 | ~ spl42_308)), inference(subsumption_resolution, [], [f14394, f1631])).
fof(f1631, plain, (~ (e10 = op1(e16, e15)) | spl42_8), inference(avatar_component_clause, [], [f1630])).
fof(f1630, plain, (spl42_8 <=> (e10 = op1(e16, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_8])])).
fof(f14394, plain, ((e10 = op1(e16, e15)) | (~ spl42_92 | ~ spl42_308)), inference(backward_demodulation, [], [f14306, f2874])).
fof(f14306, plain, ((e10 = op1(op1(e10, e15), e15)) | ~ spl42_92), inference(forward_demodulation, [], [f263, f1980])).
fof(f1980, plain, ((e10 = op1(e15, e10)) | ~ spl42_92), inference(avatar_component_clause, [], [f1978])).
fof(f1978, plain, (spl42_92 <=> (e10 = op1(e15, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_92])])).
fof(f263, plain, (e10 = op1(op1(op1(e15, e10), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f14338, plain, (spl42_180 | ~ spl42_211 | ~ spl42_326), inference(avatar_split_clause, [], [f13958, f2947, f2471, f2342])).
fof(f13958, plain, ((e14 = op1(e13, e12)) | (~ spl42_211 | ~ spl42_326)), inference(backward_demodulation, [], [f13086, f2949])).
fof(f2949, plain, ((e13 = op1(e10, e12)) | ~ spl42_326), inference(avatar_component_clause, [], [f2947])).
fof(f14292, plain, (~ spl42_58 | ~ spl42_83 | spl42_255), inference(avatar_contradiction_clause, [], [f14291])).
fof(f14291, plain, ($false | (~ spl42_58 | ~ spl42_83 | spl42_255)), inference(subsumption_resolution, [], [f14290, f2654])).
fof(f2654, plain, (~ (e12 = op1(e11, e15)) | spl42_255), inference(avatar_component_clause, [], [f2653])).
fof(f14290, plain, ((e12 = op1(e11, e15)) | (~ spl42_58 | ~ spl42_83)), inference(forward_demodulation, [], [f14289, f1839])).
fof(f14289, plain, ((e12 = op1(op1(e15, e15), e15)) | ~ spl42_83), inference(forward_demodulation, [], [f277, f1942])).
fof(f1942, plain, ((e15 = op1(e15, e12)) | ~ spl42_83), inference(avatar_component_clause, [], [f1940])).
fof(f1940, plain, (spl42_83 <=> (e15 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_83])])).
fof(f14262, plain, (spl42_256 | ~ spl42_58 | ~ spl42_76), inference(avatar_split_clause, [], [f14261, f1911, f1837, f2657])).
fof(f2657, plain, (spl42_256 <=> (e13 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_256])])).
fof(f1911, plain, (spl42_76 <=> (e15 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_76])])).
fof(f14261, plain, ((e13 = op1(e11, e15)) | (~ spl42_58 | ~ spl42_76)), inference(forward_demodulation, [], [f14260, f1839])).
fof(f14260, plain, ((e13 = op1(op1(e15, e15), e15)) | ~ spl42_76), inference(forward_demodulation, [], [f284, f1913])).
fof(f1913, plain, ((e15 = op1(e15, e13)) | ~ spl42_76), inference(avatar_component_clause, [], [f1911])).
fof(f14243, plain, (~ spl42_77 | ~ spl42_56), inference(avatar_split_clause, [], [f14242, f1828, f1915])).
fof(f1915, plain, (spl42_77 <=> (e16 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_77])])).
fof(f14242, plain, (~ (e16 = op1(e15, e13)) | ~ spl42_56), inference(forward_demodulation, [], [f772, f1830])).
fof(f772, plain, ~ (op1(e15, e13) = op1(e15, e16)), inference(cnf_transformation, [], [f7])).
fof(f14213, plain, (spl42_222 | ~ spl42_24 | ~ spl42_168), inference(avatar_split_clause, [], [f14212, f2292, f1696, f2516])).
fof(f1696, plain, (spl42_24 <=> (e12 = op1(e16, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_24])])).
fof(f14212, plain, ((e14 = op1(e12, e13)) | (~ spl42_24 | ~ spl42_168)), inference(forward_demodulation, [], [f10201, f1698])).
fof(f1698, plain, ((e12 = op1(e16, e13)) | ~ spl42_24), inference(avatar_component_clause, [], [f1696])).
fof(f14199, plain, (spl42_56 | ~ spl42_14), inference(avatar_split_clause, [], [f14198, f1654, f1828])).
fof(f14198, plain, ((e16 = op1(e15, e16)) | ~ spl42_14), inference(backward_demodulation, [], [f306, f14197])).
fof(f14197, plain, ((e15 = op1(op1(e16, e16), e16)) | ~ spl42_14), inference(forward_demodulation, [], [f299, f1656])).
fof(f299, plain, (e15 = op1(op1(op1(e16, e15), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f306, plain, (e16 = op1(op1(op1(e16, e16), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f14132, plain, (~ spl42_27 | ~ spl42_48), inference(avatar_split_clause, [], [f14126, f1795, f1708])).
fof(f1708, plain, (spl42_27 <=> (e15 = op1(e16, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_27])])).
fof(f14126, plain, (~ (e15 = op1(e16, e13)) | ~ spl42_48), inference(backward_demodulation, [], [f778, f1797])).
fof(f778, plain, ~ (op1(e16, e10) = op1(e16, e13)), inference(cnf_transformation, [], [f7])).
fof(f14093, plain, (~ spl42_89 | spl42_107 | ~ spl42_110), inference(avatar_contradiction_clause, [], [f14092])).
fof(f14092, plain, ($false | (~ spl42_89 | spl42_107 | ~ spl42_110)), inference(subsumption_resolution, [], [f14090, f2041])).
fof(f2041, plain, (~ (e11 = op1(e14, e15)) | spl42_107), inference(avatar_component_clause, [], [f2040])).
fof(f2040, plain, (spl42_107 <=> (e11 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_107])])).
fof(f14090, plain, ((e11 = op1(e14, e15)) | (~ spl42_89 | ~ spl42_110)), inference(backward_demodulation, [], [f13300, f2054])).
fof(f2054, plain, ((e14 = op1(e14, e15)) | ~ spl42_110), inference(avatar_component_clause, [], [f2052])).
fof(f2052, plain, (spl42_110 <=> (e14 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_110])])).
fof(f14086, plain, (~ spl42_114 | ~ spl42_115), inference(avatar_contradiction_clause, [], [f14085])).
fof(f14085, plain, ($false | (~ spl42_114 | ~ spl42_115)), inference(subsumption_resolution, [], [f14084, f1097])).
fof(f14084, plain, ((e11 = e12) | (~ spl42_114 | ~ spl42_115)), inference(backward_demodulation, [], [f2075, f2071])).
fof(f2071, plain, ((e11 = op1(e14, e14)) | ~ spl42_114), inference(avatar_component_clause, [], [f2069])).
fof(f2069, plain, (spl42_114 <=> (e11 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_114])])).
fof(f14083, plain, (~ spl42_115 | ~ spl42_116), inference(avatar_contradiction_clause, [], [f14082])).
fof(f14082, plain, ($false | (~ spl42_115 | ~ spl42_116)), inference(subsumption_resolution, [], [f14081, f1102])).
fof(f14081, plain, ((e12 = e13) | (~ spl42_115 | ~ spl42_116)), inference(forward_demodulation, [], [f2079, f2075])).
fof(f2079, plain, ((e13 = op1(e14, e14)) | ~ spl42_116), inference(avatar_component_clause, [], [f2077])).
fof(f2077, plain, (spl42_116 <=> (e13 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_116])])).
fof(f14055, plain, (~ spl42_167 | ~ spl42_168), inference(avatar_contradiction_clause, [], [f14054])).
fof(f14054, plain, ($false | (~ spl42_167 | ~ spl42_168)), inference(subsumption_resolution, [], [f14053, f1111])).
fof(f1111, plain, ~ (e15 = e16), inference(cnf_transformation, [], [f9])).
fof(f14053, plain, ((e15 = e16) | (~ spl42_167 | ~ spl42_168)), inference(backward_demodulation, [], [f2294, f2290])).
fof(f2290, plain, ((e15 = op1(e13, e14)) | ~ spl42_167), inference(avatar_component_clause, [], [f2288])).
fof(f13935, plain, (~ spl42_374 | ~ spl42_376), inference(avatar_contradiction_clause, [], [f13934])).
fof(f13934, plain, ($false | (~ spl42_374 | ~ spl42_376)), inference(subsumption_resolution, [], [f13933, f1124])).
fof(f13933, plain, ((e22 = e24) | (~ spl42_374 | ~ spl42_376)), inference(backward_demodulation, [], [f3254, f3246])).
fof(f3246, plain, ((e22 = op2(e26, e22)) | ~ spl42_374), inference(avatar_component_clause, [], [f3244])).
fof(f3244, plain, (spl42_374 <=> (e22 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_374])])).
fof(f13890, plain, (~ spl42_412 | ~ spl42_419), inference(avatar_split_clause, [], [f13882, f3430, f3401])).
fof(f3430, plain, (spl42_419 <=> (e25 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_419])])).
fof(f13882, plain, (~ (e25 = op2(e25, e24)) | ~ spl42_419), inference(backward_demodulation, [], [f1064, f3432])).
fof(f3432, plain, ((e25 = op2(e25, e23)) | ~ spl42_419), inference(avatar_component_clause, [], [f3430])).
fof(f1064, plain, ~ (op2(e25, e23) = op2(e25, e24)), inference(cnf_transformation, [], [f8])).
fof(f13879, plain, (~ spl42_376 | ~ spl42_443 | ~ spl42_589), inference(avatar_contradiction_clause, [], [f13878])).
fof(f13878, plain, ($false | (~ spl42_376 | ~ spl42_443 | ~ spl42_589)), inference(subsumption_resolution, [], [f13877, f1113])).
fof(f13877, plain, ((e20 = e22) | (~ spl42_376 | ~ spl42_443 | ~ spl42_589)), inference(forward_demodulation, [], [f13870, f4137])).
fof(f4137, plain, ((e20 = op2(e21, e26)) | ~ spl42_589), inference(avatar_component_clause, [], [f4135])).
fof(f4135, plain, (spl42_589 <=> (e20 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_589])])).
fof(f13870, plain, ((e22 = op2(e21, e26)) | (~ spl42_376 | ~ spl42_443)), inference(backward_demodulation, [], [f12970, f3532])).
fof(f12970, plain, ((e22 = op2(op2(e24, e26), e26)) | ~ spl42_376), inference(backward_demodulation, [], [f474, f3254])).
fof(f13849, plain, (~ spl42_456 | ~ spl42_458), inference(avatar_contradiction_clause, [], [f13848])).
fof(f13848, plain, ($false | (~ spl42_456 | ~ spl42_458)), inference(subsumption_resolution, [], [f13847, f1113])).
fof(f13847, plain, ((e20 = e22) | (~ spl42_456 | ~ spl42_458)), inference(backward_demodulation, [], [f3594, f3586])).
fof(f3586, plain, ((e20 = op2(e24, e24)) | ~ spl42_456), inference(avatar_component_clause, [], [f3584])).
fof(f3584, plain, (spl42_456 <=> (e20 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_456])])).
fof(f13840, plain, (~ spl42_458 | ~ spl42_467 | ~ spl42_554), inference(avatar_contradiction_clause, [], [f13839])).
fof(f13839, plain, ($false | (~ spl42_458 | ~ spl42_467 | ~ spl42_554)), inference(subsumption_resolution, [], [f13838, f1114])).
fof(f13838, plain, ((e20 = e23) | (~ spl42_458 | ~ spl42_467 | ~ spl42_554)), inference(forward_demodulation, [], [f13837, f3992])).
fof(f13837, plain, ((e23 = op2(e22, e24)) | (~ spl42_458 | ~ spl42_467)), inference(forward_demodulation, [], [f13830, f3594])).
fof(f13830, plain, ((e23 = op2(op2(e24, e24), e24)) | ~ spl42_467), inference(backward_demodulation, [], [f479, f3631])).
fof(f3631, plain, ((e24 = op2(e24, e23)) | ~ spl42_467), inference(avatar_component_clause, [], [f3629])).
fof(f13799, plain, (~ spl42_510 | ~ spl42_511), inference(avatar_contradiction_clause, [], [f13798])).
fof(f13798, plain, ($false | (~ spl42_510 | ~ spl42_511)), inference(subsumption_resolution, [], [f13797, f1132])).
fof(f13797, plain, ((e25 = e26) | (~ spl42_510 | ~ spl42_511)), inference(backward_demodulation, [], [f3813, f3809])).
fof(f3809, plain, ((e25 = op2(e23, e24)) | ~ spl42_510), inference(avatar_component_clause, [], [f3807])).
fof(f13796, plain, (~ spl42_516 | spl42_801), inference(avatar_contradiction_clause, [], [f13795])).
fof(f13795, plain, ($false | (~ spl42_516 | spl42_801)), inference(subsumption_resolution, [], [f13794, f5451])).
fof(f5451, plain, (~ (e24 = h4(e11)) | spl42_801), inference(avatar_component_clause, [], [f5449])).
fof(f13794, plain, ((e24 = h4(e11)) | ~ spl42_516), inference(backward_demodulation, [], [f1231, f3834])).
fof(f13775, plain, (~ spl42_572 | spl42_843), inference(avatar_contradiction_clause, [], [f13774])).
fof(f13774, plain, ($false | (~ spl42_572 | spl42_843)), inference(subsumption_resolution, [], [f13773, f5662])).
fof(f5662, plain, (~ (e24 = h3(e11)) | spl42_843), inference(avatar_component_clause, [], [f5660])).
fof(f13773, plain, ((e24 = h3(e11)) | ~ spl42_572), inference(backward_demodulation, [], [f1224, f4066])).
fof(f4066, plain, ((e24 = op2(e22, e22)) | ~ spl42_572), inference(avatar_component_clause, [], [f4064])).
fof(f4064, plain, (spl42_572 <=> (e24 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_572])])).
fof(f13761, plain, (~ spl42_432 | ~ spl42_480 | ~ spl42_594), inference(avatar_contradiction_clause, [], [f13760])).
fof(f13760, plain, ($false | (~ spl42_432 | ~ spl42_480 | ~ spl42_594)), inference(subsumption_resolution, [], [f13759, f1129])).
fof(f13759, plain, ((e23 = e26) | (~ spl42_432 | ~ spl42_480 | ~ spl42_594)), inference(forward_demodulation, [], [f13758, f3685])).
fof(f13758, plain, ((e26 = op2(e24, e21)) | (~ spl42_432 | ~ spl42_594)), inference(forward_demodulation, [], [f13751, f3486])).
fof(f13751, plain, ((e26 = op2(op2(e25, e21), e21)) | ~ spl42_594), inference(backward_demodulation, [], [f497, f4157])).
fof(f4157, plain, ((e25 = op2(e21, e26)) | ~ spl42_594), inference(avatar_component_clause, [], [f4155])).
fof(f4155, plain, (spl42_594 <=> (e25 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_594])])).
fof(f13694, plain, (~ spl42_684 | spl42_927), inference(avatar_contradiction_clause, [], [f13693])).
fof(f13693, plain, ($false | (~ spl42_684 | spl42_927)), inference(subsumption_resolution, [], [f13692, f6084])).
fof(f6084, plain, (~ (e24 = h1(e11)) | spl42_927), inference(avatar_component_clause, [], [f6082])).
fof(f6082, plain, (spl42_927 <=> (e24 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_927])])).
fof(f13692, plain, ((e24 = h1(e11)) | ~ spl42_684), inference(backward_demodulation, [], [f1210, f4530])).
fof(f4530, plain, ((op2(e20, e20) = e24) | ~ spl42_684), inference(avatar_component_clause, [], [f4528])).
fof(f4528, plain, (spl42_684 <=> (op2(e20, e20) = e24)), introduced(avatar_definition, [new_symbols(naming, [spl42_684])])).
fof(f13674, plain, (spl42_619 | ~ spl42_577), inference(avatar_split_clause, [], [f12819, f4085, f4259])).
fof(f4085, plain, (spl42_577 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_577])])).
fof(f12819, plain, ((e22 = op2(e21, e22)) | ~ spl42_577), inference(backward_demodulation, [], [f4721, f12818])).
fof(f12818, plain, ((e21 = op2(h3(e11), e22)) | ~ spl42_577), inference(forward_demodulation, [], [f12811, f1224])).
fof(f12811, plain, ((e21 = op2(op2(e22, e22), e22)) | ~ spl42_577), inference(backward_demodulation, [], [f463, f4087])).
fof(f4087, plain, ((e22 = op2(e22, e21)) | ~ spl42_577), inference(avatar_component_clause, [], [f4085])).
fof(f13671, plain, (~ spl42_850 | ~ spl42_550), inference(avatar_split_clause, [], [f13670, f3973, f5695])).
fof(f5695, plain, (spl42_850 <=> (e23 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_850])])).
fof(f13670, plain, (~ (e23 = h3(e11)) | ~ spl42_550), inference(forward_demodulation, [], [f4732, f3975])).
fof(f4732, plain, ~ (op2(e22, e25) = h3(e11)), inference(backward_demodulation, [], [f999, f1224])).
fof(f999, plain, ~ (op2(e22, e22) = op2(e22, e25)), inference(cnf_transformation, [], [f8])).
fof(f13663, plain, (~ spl42_564 | ~ spl42_550), inference(avatar_split_clause, [], [f13662, f3973, f4031])).
fof(f4031, plain, (spl42_564 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_564])])).
fof(f13662, plain, (~ (e23 = op2(e22, e23)) | ~ spl42_550), inference(forward_demodulation, [], [f1002, f3975])).
fof(f1002, plain, ~ (op2(e22, e23) = op2(e22, e25)), inference(cnf_transformation, [], [f8])).
fof(f13623, plain, (~ spl42_501 | ~ spl42_550), inference(avatar_split_clause, [], [f13622, f3973, f3770])).
fof(f13622, plain, (~ (e23 = op2(e23, e25)) | ~ spl42_550), inference(forward_demodulation, [], [f913, f3975])).
fof(f913, plain, ~ (op2(e22, e25) = op2(e23, e25)), inference(cnf_transformation, [], [f8])).
fof(f13587, plain, (~ spl42_444 | ~ spl42_743), inference(avatar_split_clause, [], [f12622, f5142, f3534])).
fof(f3534, plain, (spl42_444 <=> (e22 = op2(e24, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_444])])).
fof(f12622, plain, (~ (e22 = op2(e24, e26)) | ~ spl42_743), inference(forward_demodulation, [], [f4776, f5143])).
fof(f4776, plain, ~ (op2(e24, e26) = h5(e11)), inference(backward_demodulation, [], [f1047, f1238])).
fof(f1047, plain, ~ (op2(e24, e24) = op2(e24, e26)), inference(cnf_transformation, [], [f8])).
fof(f13551, plain, (~ spl42_418 | ~ spl42_432), inference(avatar_split_clause, [], [f13550, f3484, f3426])).
fof(f3426, plain, (spl42_418 <=> (e24 = op2(e25, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_418])])).
fof(f13550, plain, (~ (e24 = op2(e25, e23)) | ~ spl42_432), inference(forward_demodulation, [], [f1056, f3486])).
fof(f1056, plain, ~ (op2(e25, e21) = op2(e25, e23)), inference(cnf_transformation, [], [f8])).
fof(f13545, plain, (~ spl42_413 | ~ spl42_1016), inference(avatar_split_clause, [], [f12566, f6684, f3405])).
fof(f3405, plain, (spl42_413 <=> (e26 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_413])])).
fof(f12566, plain, (~ (e26 = op2(e25, e24)) | ~ spl42_1016), inference(forward_demodulation, [], [f4790, f6685])).
fof(f4790, plain, ~ (op2(e25, e24) = h6(e16)), inference(backward_demodulation, [], [f897, f4785])).
fof(f897, plain, ~ (op2(e23, e24) = op2(e25, e24)), inference(cnf_transformation, [], [f8])).
fof(f13543, plain, (spl42_595 | ~ spl42_385), inference(avatar_split_clause, [], [f13542, f3289, f4159])).
fof(f13542, plain, ((e26 = op2(e21, e26)) | ~ spl42_385), inference(backward_demodulation, [], [f4826, f13541])).
fof(f13541, plain, ((e21 = op2(h7(e11), e26)) | ~ spl42_385), inference(forward_demodulation, [], [f13540, f1252])).
fof(f13540, plain, ((e21 = op2(op2(e26, e26), e26)) | ~ spl42_385), inference(forward_demodulation, [], [f467, f3291])).
fof(f3291, plain, ((e26 = op2(e26, e21)) | ~ spl42_385), inference(avatar_component_clause, [], [f3289])).
fof(f467, plain, (e21 = op2(op2(op2(e26, e21), e26), e26)), inference(cnf_transformation, [], [f6])).
fof(f13504, plain, (~ spl42_360 | ~ spl42_743), inference(avatar_split_clause, [], [f13503, f5142, f3186])).
fof(f3186, plain, (spl42_360 <=> (e22 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_360])])).
fof(f13503, plain, (~ (e22 = op2(e26, e24)) | ~ spl42_743), inference(forward_demodulation, [], [f4770, f5143])).
fof(f4770, plain, ~ (op2(e26, e24) = h5(e11)), inference(backward_demodulation, [], [f900, f1238])).
fof(f900, plain, ~ (op2(e24, e24) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f13499, plain, (~ spl42_363 | ~ spl42_391), inference(avatar_split_clause, [], [f13498, f3314, f3198])).
fof(f13498, plain, (~ (e25 = op2(e26, e24)) | ~ spl42_391), inference(forward_demodulation, [], [f1073, f3316])).
fof(f1073, plain, ~ (op2(e26, e20) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f13426, plain, (~ spl42_256 | ~ spl42_207), inference(avatar_split_clause, [], [f13425, f2454, f2657])).
fof(f13425, plain, (~ (e13 = op1(e11, e15)) | ~ spl42_207), inference(forward_demodulation, [], [f614, f2456])).
fof(f614, plain, ~ (op1(e11, e15) = op1(e12, e15)), inference(cnf_transformation, [], [f7])).
fof(f13404, plain, (~ spl42_221 | ~ spl42_207), inference(avatar_split_clause, [], [f13403, f2454, f2512])).
fof(f2512, plain, (spl42_221 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_221])])).
fof(f13403, plain, (~ (e13 = op1(e12, e13)) | ~ spl42_207), inference(forward_demodulation, [], [f708, f2456])).
fof(f708, plain, ~ (op1(e12, e13) = op1(e12, e15)), inference(cnf_transformation, [], [f7])).
fof(f13359, plain, (~ spl42_158 | ~ spl42_207), inference(avatar_split_clause, [], [f13358, f2454, f2251])).
fof(f2251, plain, (spl42_158 <=> (e13 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_158])])).
fof(f13358, plain, (~ (e13 = op1(e13, e15)) | ~ spl42_207), inference(forward_demodulation, [], [f619, f2456])).
fof(f619, plain, ~ (op1(e12, e15) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f13336, plain, (~ spl42_82 | ~ spl42_89), inference(avatar_contradiction_clause, [], [f13335])).
fof(f13335, plain, ($false | (~ spl42_82 | ~ spl42_89)), inference(subsumption_resolution, [], [f13334, f1097])).
fof(f13334, plain, ((e11 = e12) | (~ spl42_82 | ~ spl42_89)), inference(forward_demodulation, [], [f12271, f13300])).
fof(f12271, plain, ((e12 = op1(op1(e14, e15), e15)) | ~ spl42_82), inference(forward_demodulation, [], [f277, f1938])).
fof(f1938, plain, ((e14 = op1(e15, e12)) | ~ spl42_82), inference(avatar_component_clause, [], [f1936])).
fof(f1936, plain, (spl42_82 <=> (e14 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_82])])).
fof(f13331, plain, (~ spl42_109 | ~ spl42_207), inference(avatar_split_clause, [], [f13330, f2454, f2048])).
fof(f2048, plain, (spl42_109 <=> (e13 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_109])])).
fof(f13330, plain, (~ (e13 = op1(e14, e15)) | ~ spl42_207), inference(forward_demodulation, [], [f620, f2456])).
fof(f620, plain, ~ (op1(e12, e15) = op1(e14, e15)), inference(cnf_transformation, [], [f7])).
fof(f13311, plain, (~ spl42_103 | ~ spl42_115 | ~ spl42_211), inference(avatar_contradiction_clause, [], [f13310])).
fof(f13310, plain, ($false | (~ spl42_103 | ~ spl42_115 | ~ spl42_211)), inference(subsumption_resolution, [], [f13309, f1096])).
fof(f1096, plain, ~ (e10 = e16), inference(cnf_transformation, [], [f9])).
fof(f13309, plain, ((e10 = e16) | (~ spl42_103 | ~ spl42_115 | ~ spl42_211)), inference(forward_demodulation, [], [f13308, f2473])).
fof(f13308, plain, ((e16 = op1(e12, e14)) | (~ spl42_103 | ~ spl42_115)), inference(forward_demodulation, [], [f13307, f2075])).
fof(f13307, plain, ((e16 = op1(op1(e14, e14), e14)) | ~ spl42_103), inference(forward_demodulation, [], [f304, f2025])).
fof(f2025, plain, ((e14 = op1(e14, e16)) | ~ spl42_103), inference(avatar_component_clause, [], [f2023])).
fof(f2023, plain, (spl42_103 <=> (e14 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_103])])).
fof(f13299, plain, (~ spl42_138 | ~ spl42_89), inference(avatar_split_clause, [], [f13298, f1965, f2168])).
fof(f2168, plain, (spl42_138 <=> (e14 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_138])])).
fof(f13298, plain, (~ (e14 = op1(e14, e11)) | ~ spl42_89), inference(forward_demodulation, [], [f542, f1967])).
fof(f542, plain, ~ (op1(e14, e11) = op1(e15, e11)), inference(cnf_transformation, [], [f7])).
fof(f13295, plain, (spl42_138 | ~ spl42_89 | ~ spl42_265), inference(avatar_split_clause, [], [f13294, f2694, f1965, f2168])).
fof(f2694, plain, (spl42_265 <=> (e15 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_265])])).
fof(f13294, plain, ((e14 = op1(e14, e11)) | (~ spl42_89 | ~ spl42_265)), inference(forward_demodulation, [], [f13033, f1967])).
fof(f13033, plain, ((e14 = op1(op1(e15, e11), e11)) | ~ spl42_265), inference(backward_demodulation, [], [f287, f2696])).
fof(f2696, plain, ((e15 = op1(e11, e14)) | ~ spl42_265), inference(avatar_component_clause, [], [f2694])).
fof(f287, plain, (e14 = op1(op1(op1(e11, e14), e11), e11)), inference(cnf_transformation, [], [f3])).
fof(f13282, plain, (spl42_158 | ~ spl42_73 | ~ spl42_207), inference(avatar_split_clause, [], [f13281, f2454, f1899, f2251])).
fof(f1899, plain, (spl42_73 <=> (e12 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_73])])).
fof(f13281, plain, ((e13 = op1(e13, e15)) | (~ spl42_73 | ~ spl42_207)), inference(forward_demodulation, [], [f13280, f2456])).
fof(f13280, plain, ((e13 = op1(op1(e12, e15), e15)) | ~ spl42_73), inference(forward_demodulation, [], [f284, f1901])).
fof(f1901, plain, ((e12 = op1(e15, e13)) | ~ spl42_73), inference(avatar_component_clause, [], [f1899])).
fof(f13266, plain, (spl42_161 | ~ spl42_52 | ~ spl42_207), inference(avatar_split_clause, [], [f13265, f2454, f1812, f2263])).
fof(f2263, plain, (spl42_161 <=> (e16 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_161])])).
fof(f1812, plain, (spl42_52 <=> (e12 = op1(e15, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_52])])).
fof(f13265, plain, ((e16 = op1(e13, e15)) | (~ spl42_52 | ~ spl42_207)), inference(forward_demodulation, [], [f13264, f2456])).
fof(f13264, plain, ((e16 = op1(op1(e12, e15), e15)) | ~ spl42_52), inference(forward_demodulation, [], [f305, f1814])).
fof(f1814, plain, ((e12 = op1(e15, e16)) | ~ spl42_52), inference(avatar_component_clause, [], [f1812])).
fof(f13232, plain, (spl42_24 | ~ spl42_178 | ~ spl42_224), inference(avatar_split_clause, [], [f13124, f2524, f2334, f1696])).
fof(f2524, plain, (spl42_224 <=> (e16 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_224])])).
fof(f13124, plain, ((e12 = op1(e16, e13)) | (~ spl42_178 | ~ spl42_224)), inference(forward_demodulation, [], [f13121, f2526])).
fof(f2526, plain, ((e16 = op1(e12, e13)) | ~ spl42_224), inference(avatar_component_clause, [], [f2524])).
fof(f13121, plain, ((e12 = op1(op1(e12, e13), e13)) | ~ spl42_178), inference(backward_demodulation, [], [f275, f2336])).
fof(f275, plain, (e12 = op1(op1(op1(e13, e12), e13), e13)), inference(cnf_transformation, [], [f3])).
fof(f13185, plain, (~ spl42_57 | ~ spl42_58), inference(avatar_contradiction_clause, [], [f13184])).
fof(f13184, plain, ($false | (~ spl42_57 | ~ spl42_58)), inference(subsumption_resolution, [], [f13183, f1091])).
fof(f13183, plain, ((e10 = e11) | (~ spl42_57 | ~ spl42_58)), inference(backward_demodulation, [], [f1839, f1835])).
fof(f1835, plain, ((e10 = op1(e15, e15)) | ~ spl42_57), inference(avatar_component_clause, [], [f1833])).
fof(f1833, plain, (spl42_57 <=> (e10 = op1(e15, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_57])])).
fof(f13131, plain, (~ spl42_163 | ~ spl42_168), inference(avatar_contradiction_clause, [], [f13130])).
fof(f13130, plain, ($false | (~ spl42_163 | ~ spl42_168)), inference(subsumption_resolution, [], [f13129, f1101])).
fof(f13129, plain, ((e11 = e16) | (~ spl42_163 | ~ spl42_168)), inference(backward_demodulation, [], [f2294, f2274])).
fof(f2274, plain, ((e11 = op1(e13, e14)) | ~ spl42_163), inference(avatar_component_clause, [], [f2272])).
fof(f2272, plain, (spl42_163 <=> (e11 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_163])])).
fof(f13119, plain, (~ spl42_137 | ~ spl42_186), inference(avatar_split_clause, [], [f13114, f2367, f2164])).
fof(f2367, plain, (spl42_186 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_186])])).
fof(f13114, plain, (~ (e13 = op1(e14, e11)) | ~ spl42_186), inference(backward_demodulation, [], [f539, f2369])).
fof(f2369, plain, ((e13 = op1(e13, e11)) | ~ spl42_186), inference(avatar_component_clause, [], [f2367])).
fof(f539, plain, ~ (op1(e13, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f13096, plain, (~ spl42_207 | ~ spl42_208), inference(avatar_contradiction_clause, [], [f13095])).
fof(f13095, plain, ($false | (~ spl42_207 | ~ spl42_208)), inference(subsumption_resolution, [], [f13094, f1106])).
fof(f13094, plain, ((e13 = e14) | (~ spl42_207 | ~ spl42_208)), inference(backward_demodulation, [], [f2460, f2456])).
fof(f2460, plain, ((e14 = op1(e12, e15)) | ~ spl42_208), inference(avatar_component_clause, [], [f2458])).
fof(f2458, plain, (spl42_208 <=> (e14 = op1(e12, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_208])])).
fof(f12947, plain, (~ spl42_400 | ~ spl42_401), inference(avatar_contradiction_clause, [], [f12946])).
fof(f12946, plain, ($false | (~ spl42_400 | ~ spl42_401)), inference(subsumption_resolution, [], [f12945, f1112])).
fof(f12945, plain, ((e20 = e21) | (~ spl42_400 | ~ spl42_401)), inference(backward_demodulation, [], [f3358, f3354])).
fof(f3354, plain, ((e20 = op2(e25, e25)) | ~ spl42_400), inference(avatar_component_clause, [], [f3352])).
fof(f3352, plain, (spl42_400 <=> (e20 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_400])])).
fof(f3358, plain, ((e21 = op2(e25, e25)) | ~ spl42_401), inference(avatar_component_clause, [], [f3356])).
fof(f3356, plain, (spl42_401 <=> (e21 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_401])])).
fof(f12944, plain, (~ spl42_401 | ~ spl42_402), inference(avatar_contradiction_clause, [], [f12943])).
fof(f12943, plain, ($false | (~ spl42_401 | ~ spl42_402)), inference(subsumption_resolution, [], [f12942, f1118])).
fof(f1118, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f10])).
fof(f12942, plain, ((e21 = e22) | (~ spl42_401 | ~ spl42_402)), inference(forward_demodulation, [], [f3362, f3358])).
fof(f3362, plain, ((e22 = op2(e25, e25)) | ~ spl42_402), inference(avatar_component_clause, [], [f3360])).
fof(f3360, plain, (spl42_402 <=> (e22 = op2(e25, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_402])])).
fof(f12921, plain, (~ spl42_457 | ~ spl42_458), inference(avatar_contradiction_clause, [], [f12920])).
fof(f12920, plain, ($false | (~ spl42_457 | ~ spl42_458)), inference(subsumption_resolution, [], [f12919, f1118])).
fof(f12919, plain, ((e21 = e22) | (~ spl42_457 | ~ spl42_458)), inference(backward_demodulation, [], [f3594, f3590])).
fof(f3590, plain, ((e21 = op2(e24, e24)) | ~ spl42_457), inference(avatar_component_clause, [], [f3588])).
fof(f3588, plain, (spl42_457 <=> (e21 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_457])])).
fof(f12877, plain, (spl42_358 | ~ spl42_487 | ~ spl42_511), inference(avatar_contradiction_clause, [], [f12876])).
fof(f12876, plain, ($false | (spl42_358 | ~ spl42_487 | ~ spl42_511)), inference(subsumption_resolution, [], [f12875, f3179])).
fof(f3179, plain, (~ (e20 = op2(e26, e24)) | spl42_358), inference(avatar_component_clause, [], [f3178])).
fof(f3178, plain, (spl42_358 <=> (e20 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_358])])).
fof(f12875, plain, ((e20 = op2(e26, e24)) | (~ spl42_487 | ~ spl42_511)), inference(forward_demodulation, [], [f12868, f3813])).
fof(f12868, plain, ((e20 = op2(op2(e23, e24), e24)) | ~ spl42_487), inference(backward_demodulation, [], [f458, f3714])).
fof(f3714, plain, ((e23 = op2(e24, e20)) | ~ spl42_487), inference(avatar_component_clause, [], [f3712])).
fof(f3712, plain, (spl42_487 <=> (e23 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_487])])).
fof(f458, plain, (e20 = op2(op2(op2(e24, e20), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f12849, plain, (~ spl42_550 | ~ spl42_551), inference(avatar_contradiction_clause, [], [f12848])).
fof(f12848, plain, ($false | (~ spl42_550 | ~ spl42_551)), inference(subsumption_resolution, [], [f12847, f1127])).
fof(f12847, plain, ((e23 = e24) | (~ spl42_550 | ~ spl42_551)), inference(backward_demodulation, [], [f3979, f3975])).
fof(f3979, plain, ((e24 = op2(e22, e25)) | ~ spl42_551), inference(avatar_component_clause, [], [f3977])).
fof(f3977, plain, (spl42_551 <=> (e24 = op2(e22, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_551])])).
fof(f12832, plain, (~ spl42_469 | ~ spl42_567), inference(avatar_split_clause, [], [f12828, f4043, f3637])).
fof(f4043, plain, (spl42_567 <=> (e26 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_567])])).
fof(f12828, plain, (~ (e26 = op2(e24, e23)) | ~ spl42_567), inference(backward_demodulation, [], [f872, f4045])).
fof(f4045, plain, ((e26 = op2(e22, e23)) | ~ spl42_567), inference(avatar_component_clause, [], [f4043])).
fof(f872, plain, ~ (op2(e22, e23) = op2(e24, e23)), inference(cnf_transformation, [], [f8])).
fof(f12826, plain, (spl42_850 | ~ spl42_571), inference(avatar_split_clause, [], [f12825, f4060, f5695])).
fof(f4060, plain, (spl42_571 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_571])])).
fof(f12825, plain, ((e23 = h3(e11)) | ~ spl42_571), inference(backward_demodulation, [], [f1224, f4062])).
fof(f4062, plain, ((e23 = op2(e22, e22)) | ~ spl42_571), inference(avatar_component_clause, [], [f4060])).
fof(f12800, plain, (spl42_481 | ~ spl42_432 | ~ spl42_608), inference(avatar_split_clause, [], [f12799, f4213, f3484, f3687])).
fof(f3687, plain, (spl42_481 <=> (e24 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_481])])).
fof(f4213, plain, (spl42_608 <=> (e25 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_608])])).
fof(f12799, plain, ((e24 = op2(e24, e21)) | (~ spl42_432 | ~ spl42_608)), inference(forward_demodulation, [], [f12794, f3486])).
fof(f12794, plain, ((e24 = op2(op2(e25, e21), e21)) | ~ spl42_608), inference(backward_demodulation, [], [f483, f4215])).
fof(f4215, plain, ((e25 = op2(e21, e24)) | ~ spl42_608), inference(avatar_component_clause, [], [f4213])).
fof(f12793, plain, (spl42_568 | ~ spl42_583 | ~ spl42_619), inference(avatar_contradiction_clause, [], [f12792])).
fof(f12792, plain, ($false | (spl42_568 | ~ spl42_583 | ~ spl42_619)), inference(subsumption_resolution, [], [f12787, f4049])).
fof(f4049, plain, (~ (e20 = op2(e22, e22)) | spl42_568), inference(avatar_component_clause, [], [f4048])).
fof(f4048, plain, (spl42_568 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_568])])).
fof(f12787, plain, ((e20 = op2(e22, e22)) | (~ spl42_583 | ~ spl42_619)), inference(backward_demodulation, [], [f11450, f4261])).
fof(f11450, plain, ((e20 = op2(op2(e21, e22), e22)) | ~ spl42_583), inference(backward_demodulation, [], [f456, f4112])).
fof(f4112, plain, ((e21 = op2(e22, e20)) | ~ spl42_583), inference(avatar_component_clause, [], [f4110])).
fof(f4110, plain, (spl42_583 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_583])])).
fof(f12760, plain, (spl42_934 | ~ spl42_683), inference(avatar_split_clause, [], [f12759, f4524, f6117])).
fof(f12759, plain, ((e23 = h1(e11)) | ~ spl42_683), inference(backward_demodulation, [], [f1210, f4526])).
fof(f12756, plain, (~ spl42_660 | ~ spl42_611), inference(avatar_split_clause, [], [f12755, f4226, f4429])).
fof(f4226, plain, (spl42_611 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_611])])).
fof(f12755, plain, (~ (e21 = op2(e20, e23)) | ~ spl42_611), inference(forward_demodulation, [], [f860, f4228])).
fof(f4228, plain, ((e21 = op2(e21, e23)) | ~ spl42_611), inference(avatar_component_clause, [], [f4226])).
fof(f860, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f8])).
fof(f12735, plain, (~ spl42_618 | ~ spl42_275 | ~ spl42_743 | ~ spl42_748 | spl42_1053), inference(avatar_split_clause, [], [f12734, f6832, f5171, f5142, f2736, f4255])).
fof(f4255, plain, (spl42_618 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_618])])).
fof(f2736, plain, (spl42_275 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_275])])).
fof(f6832, plain, (spl42_1053 <=> (h6(op1(e11, e12)) = op2(h6(e11), h5(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl42_1053])])).
fof(f12734, plain, (~ (e21 = op2(e21, e22)) | (~ spl42_275 | ~ spl42_743 | ~ spl42_748 | spl42_1053)), inference(forward_demodulation, [], [f12733, f5172])).
fof(f12733, plain, (~ (op2(e21, e22) = h6(e11)) | (~ spl42_275 | ~ spl42_743 | ~ spl42_748 | spl42_1053)), inference(forward_demodulation, [], [f11299, f2738])).
fof(f2738, plain, ((e11 = op1(e11, e12)) | ~ spl42_275), inference(avatar_component_clause, [], [f2736])).
fof(f11299, plain, (~ (op2(e21, e22) = h6(op1(e11, e12))) | (~ spl42_743 | ~ spl42_748 | spl42_1053)), inference(backward_demodulation, [], [f11275, f5143])).
fof(f11275, plain, (~ (h6(op1(e11, e12)) = op2(e21, h5(e11))) | (~ spl42_748 | spl42_1053)), inference(backward_demodulation, [], [f6834, f5172])).
fof(f6834, plain, (~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | spl42_1053), inference(avatar_component_clause, [], [f6832])).
fof(f12725, plain, (~ spl42_579 | ~ spl42_432), inference(avatar_split_clause, [], [f12724, f3484, f4093])).
fof(f4093, plain, (spl42_579 <=> (e24 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_579])])).
fof(f12724, plain, (~ (e24 = op2(e22, e21)) | ~ spl42_432), inference(forward_demodulation, [], [f831, f3486])).
fof(f831, plain, ~ (op2(e22, e21) = op2(e25, e21)), inference(cnf_transformation, [], [f8])).
fof(f12665, plain, (~ spl42_494 | ~ spl42_496), inference(avatar_contradiction_clause, [], [f12664])).
fof(f12664, plain, ($false | (~ spl42_494 | ~ spl42_496)), inference(subsumption_resolution, [], [f12663, f1128])).
fof(f12663, plain, ((e23 = e25) | (~ spl42_494 | ~ spl42_496)), inference(backward_demodulation, [], [f3751, f3743])).
fof(f3751, plain, ((e25 = op2(e23, e26)) | ~ spl42_496), inference(avatar_component_clause, [], [f3749])).
fof(f3749, plain, (spl42_496 <=> (e25 = op2(e23, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_496])])).
fof(f12644, plain, (~ spl42_481 | ~ spl42_432), inference(avatar_split_clause, [], [f12643, f3484, f3687])).
fof(f12643, plain, (~ (e24 = op2(e24, e21)) | ~ spl42_432), inference(forward_demodulation, [], [f836, f3486])).
fof(f836, plain, ~ (op2(e24, e21) = op2(e25, e21)), inference(cnf_transformation, [], [f8])).
fof(f12621, plain, (~ spl42_442 | ~ spl42_458 | ~ spl42_656), inference(avatar_contradiction_clause, [], [f12620])).
fof(f12620, plain, ($false | (~ spl42_442 | ~ spl42_458 | ~ spl42_656)), inference(subsumption_resolution, [], [f12619, f1126])).
fof(f12619, plain, ((e22 = e26) | (~ spl42_442 | ~ spl42_458 | ~ spl42_656)), inference(forward_demodulation, [], [f12618, f3594])).
fof(f12618, plain, ((e26 = op2(e24, e24)) | (~ spl42_442 | ~ spl42_656)), inference(forward_demodulation, [], [f12617, f4414])).
fof(f12617, plain, ((e26 = op2(op2(e20, e24), e24)) | ~ spl42_442), inference(forward_demodulation, [], [f500, f3528])).
fof(f3528, plain, ((e20 = op2(e24, e26)) | ~ spl42_442), inference(avatar_component_clause, [], [f3526])).
fof(f500, plain, (e26 = op2(op2(op2(e24, e26), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f12595, plain, (spl42_499 | ~ spl42_432 | ~ spl42_452), inference(avatar_split_clause, [], [f12594, f3567, f3484, f3762])).
fof(f3762, plain, (spl42_499 <=> (e21 = op2(e23, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_499])])).
fof(f3567, plain, (spl42_452 <=> (e23 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_452])])).
fof(f12594, plain, ((e21 = op2(e23, e25)) | (~ spl42_432 | ~ spl42_452)), inference(forward_demodulation, [], [f12593, f3569])).
fof(f3569, plain, ((e23 = op2(e24, e25)) | ~ spl42_452), inference(avatar_component_clause, [], [f3567])).
fof(f12575, plain, (spl42_600 | ~ spl42_401 | ~ spl42_412), inference(avatar_split_clause, [], [f12574, f3401, f3356, f4180])).
fof(f12574, plain, ((e24 = op2(e21, e25)) | (~ spl42_401 | ~ spl42_412)), inference(forward_demodulation, [], [f12573, f3358])).
fof(f12573, plain, ((e24 = op2(op2(e25, e25), e25)) | ~ spl42_412), inference(forward_demodulation, [], [f487, f3403])).
fof(f487, plain, (e24 = op2(op2(op2(e25, e24), e25), e25)), inference(cnf_transformation, [], [f6])).
fof(f12569, plain, (~ spl42_409 | ~ spl42_743), inference(avatar_split_clause, [], [f12568, f5142, f3389])).
fof(f3389, plain, (spl42_409 <=> (e22 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_409])])).
fof(f12568, plain, (~ (e22 = op2(e25, e24)) | ~ spl42_743), inference(forward_demodulation, [], [f4769, f5143])).
fof(f4769, plain, ~ (op2(e25, e24) = h5(e11)), inference(backward_demodulation, [], [f899, f1238])).
fof(f899, plain, ~ (op2(e24, e24) = op2(e25, e24)), inference(cnf_transformation, [], [f8])).
fof(f12565, plain, (~ spl42_408 | ~ spl42_748), inference(avatar_split_clause, [], [f12564, f5171, f3385])).
fof(f3385, plain, (spl42_408 <=> (e21 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_408])])).
fof(f12564, plain, (~ (e21 = op2(e25, e24)) | ~ spl42_748), inference(forward_demodulation, [], [f4814, f5172])).
fof(f4814, plain, ~ (op2(e25, e24) = h6(e11)), inference(backward_demodulation, [], [f1067, f1245])).
fof(f1067, plain, ~ (op2(e25, e24) = op2(e25, e25)), inference(cnf_transformation, [], [f8])).
fof(f12534, plain, (~ spl42_394 | ~ spl42_748), inference(avatar_split_clause, [], [f12533, f5171, f3327])).
fof(f3327, plain, (spl42_394 <=> (e21 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_394])])).
fof(f12533, plain, (~ (e21 = op2(e25, e26)) | ~ spl42_748), inference(forward_demodulation, [], [f4815, f5172])).
fof(f4815, plain, ~ (op2(e25, e26) = h6(e11)), inference(backward_demodulation, [], [f1069, f1245])).
fof(f1069, plain, ~ (op2(e25, e25) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f12530, plain, (~ spl42_397 | ~ spl42_432), inference(avatar_split_clause, [], [f12529, f3484, f3339])).
fof(f3339, plain, (spl42_397 <=> (e24 = op2(e25, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_397])])).
fof(f12529, plain, (~ (e24 = op2(e25, e26)) | ~ spl42_432), inference(forward_demodulation, [], [f1059, f3486])).
fof(f1059, plain, ~ (op2(e25, e21) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f12491, plain, (spl42_418 | ~ spl42_370 | ~ spl42_1016), inference(avatar_split_clause, [], [f12485, f6684, f3227, f3426])).
fof(f3227, plain, (spl42_370 <=> (e25 = op2(e26, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_370])])).
fof(f12485, plain, ((e24 = op2(e25, e23)) | (~ spl42_370 | ~ spl42_1016)), inference(backward_demodulation, [], [f11251, f3229])).
fof(f3229, plain, ((e25 = op2(e26, e23)) | ~ spl42_370), inference(avatar_component_clause, [], [f3227])).
fof(f12475, plain, (~ spl42_368 | ~ spl42_361), inference(avatar_split_clause, [], [f10997, f3190, f3219])).
fof(f3190, plain, (spl42_361 <=> (e23 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_361])])).
fof(f10997, plain, (~ (e23 = op2(e26, e23)) | ~ spl42_361), inference(forward_demodulation, [], [f1085, f3192])).
fof(f3192, plain, ((e23 = op2(e26, e24)) | ~ spl42_361), inference(avatar_component_clause, [], [f3190])).
fof(f1085, plain, ~ (op2(e26, e23) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f12442, plain, (~ spl42_327 | ~ spl42_313), inference(avatar_split_clause, [], [f12441, f2893, f2951])).
fof(f2951, plain, (spl42_327 <=> (e14 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_327])])).
fof(f12441, plain, (~ (e14 = op1(e10, e12)) | ~ spl42_313), inference(forward_demodulation, [], [f662, f2895])).
fof(f662, plain, ~ (op1(e10, e12) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f12413, plain, (spl42_257 | ~ spl42_58 | ~ spl42_69), inference(avatar_split_clause, [], [f12109, f1882, f1837, f2661])).
fof(f12109, plain, ((e14 = op1(e11, e15)) | (~ spl42_58 | ~ spl42_69)), inference(forward_demodulation, [], [f12107, f1839])).
fof(f12107, plain, ((e14 = op1(op1(e15, e15), e15)) | ~ spl42_69), inference(backward_demodulation, [], [f291, f1884])).
fof(f291, plain, (e14 = op1(op1(op1(e15, e14), e15), e15)), inference(cnf_transformation, [], [f3])).
fof(f12345, plain, (~ spl42_18 | ~ spl42_151), inference(avatar_contradiction_clause, [], [f12344])).
fof(f12344, plain, ($false | (~ spl42_18 | ~ spl42_151)), inference(subsumption_resolution, [], [f12343, f1106])).
fof(f12343, plain, ((e13 = e14) | (~ spl42_18 | ~ spl42_151)), inference(forward_demodulation, [], [f12342, f2224])).
fof(f12342, plain, ((e14 = op1(e13, e16)) | (~ spl42_18 | ~ spl42_151)), inference(forward_demodulation, [], [f10634, f2224])).
fof(f10634, plain, ((e14 = op1(op1(e13, e16), e16)) | ~ spl42_18), inference(forward_demodulation, [], [f292, f1673])).
fof(f1673, plain, ((e13 = op1(e16, e14)) | ~ spl42_18), inference(avatar_component_clause, [], [f1671])).
fof(f1671, plain, (spl42_18 <=> (e13 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_18])])).
fof(f292, plain, (e14 = op1(op1(op1(e16, e14), e16), e16)), inference(cnf_transformation, [], [f3])).
fof(f12336, plain, (~ spl42_188 | ~ spl42_153), inference(avatar_split_clause, [], [f12330, f2230, f2375])).
fof(f2230, plain, (spl42_153 <=> (e15 = op1(e13, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_153])])).
fof(f12330, plain, (~ (e15 = op1(e13, e11)) | ~ spl42_153), inference(backward_demodulation, [], [f723, f2232])).
fof(f2232, plain, ((e15 = op1(e13, e16)) | ~ spl42_153), inference(avatar_component_clause, [], [f2230])).
fof(f723, plain, ~ (op1(e13, e11) = op1(e13, e16)), inference(cnf_transformation, [], [f7])).
fof(f12299, plain, (~ spl42_99 | ~ spl42_115 | ~ spl42_313), inference(avatar_contradiction_clause, [], [f12298])).
fof(f12298, plain, ($false | (~ spl42_99 | ~ spl42_115 | ~ spl42_313)), inference(subsumption_resolution, [], [f12297, f1105])).
fof(f1105, plain, ~ (e12 = e16), inference(cnf_transformation, [], [f9])).
fof(f12297, plain, ((e12 = e16) | (~ spl42_99 | ~ spl42_115 | ~ spl42_313)), inference(forward_demodulation, [], [f12296, f2075])).
fof(f12296, plain, ((e16 = op1(e14, e14)) | (~ spl42_99 | ~ spl42_313)), inference(forward_demodulation, [], [f12295, f2895])).
fof(f12295, plain, ((e16 = op1(op1(e10, e14), e14)) | ~ spl42_99), inference(forward_demodulation, [], [f304, f2009])).
fof(f2009, plain, ((e10 = op1(e14, e16)) | ~ spl42_99), inference(avatar_component_clause, [], [f2007])).
fof(f2007, plain, (spl42_99 <=> (e10 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_99])])).
fof(f12141, plain, (spl42_13 | ~ spl42_58 | ~ spl42_259), inference(avatar_split_clause, [], [f11848, f2669, f1837, f1650])).
fof(f2669, plain, (spl42_259 <=> (e16 = op1(e11, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_259])])).
fof(f11848, plain, ((e15 = op1(e16, e15)) | (~ spl42_58 | ~ spl42_259)), inference(backward_demodulation, [], [f10430, f2671])).
fof(f2671, plain, ((e16 = op1(e11, e15)) | ~ spl42_259), inference(avatar_component_clause, [], [f2669])).
fof(f12099, plain, (spl42_305 | ~ spl42_8 | ~ spl42_77), inference(avatar_split_clause, [], [f12098, f1915, f1630, f2860])).
fof(f12098, plain, ((e13 = op1(e10, e15)) | (~ spl42_8 | ~ spl42_77)), inference(forward_demodulation, [], [f12095, f1632])).
fof(f1632, plain, ((e10 = op1(e16, e15)) | ~ spl42_8), inference(avatar_component_clause, [], [f1630])).
fof(f12095, plain, ((e13 = op1(op1(e16, e15), e15)) | ~ spl42_77), inference(backward_demodulation, [], [f284, f1917])).
fof(f1917, plain, ((e16 = op1(e15, e13)) | ~ spl42_77), inference(avatar_component_clause, [], [f1915])).
fof(f11878, plain, (~ spl42_219 | ~ spl42_233), inference(avatar_split_clause, [], [f11869, f2562, f2504])).
fof(f2504, plain, (spl42_219 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_219])])).
fof(f11869, plain, (~ (e11 = op1(e12, e13)) | ~ spl42_233), inference(backward_demodulation, [], [f699, f2564])).
fof(f699, plain, ~ (op1(e12, e11) = op1(e12, e13)), inference(cnf_transformation, [], [f7])).
fof(f11874, plain, (~ spl42_184 | ~ spl42_233), inference(avatar_split_clause, [], [f11865, f2562, f2359])).
fof(f11865, plain, (~ (e11 = op1(e13, e11)) | ~ spl42_233), inference(backward_demodulation, [], [f535, f2564])).
fof(f11843, plain, (~ spl42_100 | ~ spl42_261), inference(avatar_contradiction_clause, [], [f11842])).
fof(f11842, plain, ($false | (~ spl42_100 | ~ spl42_261)), inference(subsumption_resolution, [], [f11841, f1101])).
fof(f11841, plain, ((e11 = e16) | (~ spl42_100 | ~ spl42_261)), inference(forward_demodulation, [], [f11833, f2680])).
fof(f2680, plain, ((e11 = op1(e11, e14)) | ~ spl42_261), inference(avatar_component_clause, [], [f2678])).
fof(f2678, plain, (spl42_261 <=> (e11 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_261])])).
fof(f11833, plain, ((e16 = op1(e11, e14)) | (~ spl42_100 | ~ spl42_261)), inference(backward_demodulation, [], [f10758, f2680])).
fof(f10758, plain, ((e16 = op1(op1(e11, e14), e14)) | ~ spl42_100), inference(forward_demodulation, [], [f304, f2013])).
fof(f11737, plain, (~ spl42_304 | ~ spl42_332), inference(avatar_split_clause, [], [f11729, f2972, f2856])).
fof(f2972, plain, (spl42_332 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_332])])).
fof(f11729, plain, (~ (e12 = op1(e10, e15)) | ~ spl42_332), inference(backward_demodulation, [], [f659, f2974])).
fof(f2974, plain, ((e12 = op1(e10, e11)) | ~ spl42_332), inference(avatar_component_clause, [], [f2972])).
fof(f659, plain, ~ (op1(e10, e11) = op1(e10, e15)), inference(cnf_transformation, [], [f7])).
fof(f11610, plain, (~ spl42_464 | ~ spl42_604), inference(avatar_contradiction_clause, [], [f11609])).
fof(f11609, plain, ($false | (~ spl42_464 | ~ spl42_604)), inference(subsumption_resolution, [], [f11608, f1119])).
fof(f11608, plain, ((e21 = e23) | (~ spl42_464 | ~ spl42_604)), inference(forward_demodulation, [], [f11607, f4199])).
fof(f4199, plain, ((e21 = op2(e21, e24)) | ~ spl42_604), inference(avatar_component_clause, [], [f4197])).
fof(f11607, plain, ((e23 = op2(e21, e24)) | (~ spl42_464 | ~ spl42_604)), inference(forward_demodulation, [], [f11601, f4199])).
fof(f11601, plain, ((e23 = op2(op2(e21, e24), e24)) | ~ spl42_464), inference(backward_demodulation, [], [f479, f3619])).
fof(f3619, plain, ((e21 = op2(e24, e23)) | ~ spl42_464), inference(avatar_component_clause, [], [f3617])).
fof(f3617, plain, (spl42_464 <=> (e21 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_464])])).
fof(f11446, plain, (~ spl42_600 | ~ spl42_602), inference(avatar_contradiction_clause, [], [f11445])).
fof(f11445, plain, ($false | (~ spl42_600 | ~ spl42_602)), inference(subsumption_resolution, [], [f11444, f1131])).
fof(f11444, plain, ((e24 = e26) | (~ spl42_600 | ~ spl42_602)), inference(backward_demodulation, [], [f4190, f4182])).
fof(f4190, plain, ((e26 = op2(e21, e25)) | ~ spl42_602), inference(avatar_component_clause, [], [f4188])).
fof(f4188, plain, (spl42_602 <=> (e26 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_602])])).
fof(f11426, plain, (~ spl42_443 | ~ spl42_604), inference(avatar_contradiction_clause, [], [f11425])).
fof(f11425, plain, ($false | (~ spl42_443 | ~ spl42_604)), inference(subsumption_resolution, [], [f11424, f1122])).
fof(f1122, plain, ~ (e21 = e26), inference(cnf_transformation, [], [f10])).
fof(f11424, plain, ((e21 = e26) | (~ spl42_443 | ~ spl42_604)), inference(forward_demodulation, [], [f11418, f4199])).
fof(f11418, plain, ((e26 = op2(e21, e24)) | (~ spl42_443 | ~ spl42_604)), inference(backward_demodulation, [], [f11095, f4199])).
fof(f11095, plain, ((e26 = op2(op2(e21, e24), e24)) | ~ spl42_443), inference(forward_demodulation, [], [f500, f3532])).
fof(f11375, plain, (spl42_899 | ~ spl42_626), inference(avatar_split_clause, [], [f11374, f4288, f5941])).
fof(f4288, plain, (spl42_626 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_626])])).
fof(f11374, plain, ((e22 = h2(e11)) | ~ spl42_626), inference(backward_demodulation, [], [f1217, f4290])).
fof(f4290, plain, ((e22 = op2(e21, e21)) | ~ spl42_626), inference(avatar_component_clause, [], [f4288])).
fof(f11316, plain, (spl42_554 | ~ spl42_743), inference(avatar_split_clause, [], [f11295, f5142, f3990])).
fof(f11295, plain, ((e20 = op2(e22, e24)) | ~ spl42_743), inference(backward_demodulation, [], [f4779, f5143])).
fof(f11315, plain, (spl42_550 | ~ spl42_743), inference(avatar_split_clause, [], [f11294, f5142, f3973])).
fof(f11294, plain, ((e23 = op2(e22, e25)) | ~ spl42_743), inference(backward_demodulation, [], [f4778, f5143])).
fof(f11306, plain, (~ spl42_451 | ~ spl42_743), inference(avatar_split_clause, [], [f11292, f5142, f3563])).
fof(f3563, plain, (spl42_451 <=> (e22 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_451])])).
fof(f11292, plain, (~ (e22 = op2(e24, e25)) | ~ spl42_743), inference(backward_demodulation, [], [f4775, f5143])).
fof(f4775, plain, ~ (op2(e24, e25) = h5(e11)), inference(backward_demodulation, [], [f1046, f1238])).
fof(f1046, plain, ~ (op2(e24, e24) = op2(e24, e25)), inference(cnf_transformation, [], [f8])).
fof(f11305, plain, (~ spl42_465 | ~ spl42_743), inference(avatar_split_clause, [], [f11291, f5142, f3621])).
fof(f3621, plain, (spl42_465 <=> (e22 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_465])])).
fof(f11291, plain, (~ (e22 = op2(e24, e23)) | ~ spl42_743), inference(backward_demodulation, [], [f4774, f5143])).
fof(f4774, plain, ~ (op2(e24, e23) = h5(e11)), inference(backward_demodulation, [], [f1043, f1238])).
fof(f1043, plain, ~ (op2(e24, e23) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f11304, plain, (~ spl42_472 | ~ spl42_743), inference(avatar_split_clause, [], [f11290, f5142, f3650])).
fof(f3650, plain, (spl42_472 <=> (e22 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_472])])).
fof(f11290, plain, (~ (e22 = op2(e24, e22)) | ~ spl42_743), inference(backward_demodulation, [], [f4773, f5143])).
fof(f4773, plain, ~ (op2(e24, e22) = h5(e11)), inference(backward_demodulation, [], [f1040, f1238])).
fof(f1040, plain, ~ (op2(e24, e22) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f11303, plain, (~ spl42_479 | ~ spl42_743), inference(avatar_split_clause, [], [f11289, f5142, f3679])).
fof(f3679, plain, (spl42_479 <=> (e22 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_479])])).
fof(f11289, plain, (~ (e22 = op2(e24, e21)) | ~ spl42_743), inference(backward_demodulation, [], [f4772, f5143])).
fof(f4772, plain, ~ (op2(e24, e21) = h5(e11)), inference(backward_demodulation, [], [f1036, f1238])).
fof(f1036, plain, ~ (op2(e24, e21) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f11300, plain, (~ spl42_605 | ~ spl42_743), inference(avatar_split_clause, [], [f11286, f5142, f4201])).
fof(f4201, plain, (spl42_605 <=> (e22 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_605])])).
fof(f11286, plain, (~ (e22 = op2(e21, e24)) | ~ spl42_743), inference(backward_demodulation, [], [f4766, f5143])).
fof(f4766, plain, ~ (op2(e21, e24) = h5(e11)), inference(backward_demodulation, [], [f889, f1238])).
fof(f889, plain, ~ (op2(e21, e24) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f11285, plain, (spl42_432 | ~ spl42_748), inference(avatar_split_clause, [], [f11274, f5171, f3484])).
fof(f11274, plain, ((e24 = op2(e25, e21)) | ~ spl42_748), inference(backward_demodulation, [], [f4816, f5172])).
fof(f11284, plain, (~ spl42_415 | ~ spl42_748), inference(avatar_split_clause, [], [f11273, f5171, f3414])).
fof(f11273, plain, (~ (e21 = op2(e25, e23)) | ~ spl42_748), inference(backward_demodulation, [], [f4813, f5172])).
fof(f4813, plain, ~ (op2(e25, e23) = h6(e11)), inference(backward_demodulation, [], [f1065, f1245])).
fof(f1065, plain, ~ (op2(e25, e23) = op2(e25, e25)), inference(cnf_transformation, [], [f8])).
fof(f11280, plain, (~ spl42_450 | ~ spl42_748), inference(avatar_split_clause, [], [f11269, f5171, f3559])).
fof(f11269, plain, (~ (e21 = op2(e24, e25)) | ~ spl42_748), inference(backward_demodulation, [], [f4808, f5172])).
fof(f4808, plain, ~ (op2(e24, e25) = h6(e11)), inference(backward_demodulation, [], [f920, f1245])).
fof(f920, plain, ~ (op2(e24, e25) = op2(e25, e25)), inference(cnf_transformation, [], [f8])).
fof(f11279, plain, (~ spl42_499 | ~ spl42_748), inference(avatar_split_clause, [], [f11268, f5171, f3762])).
fof(f11268, plain, (~ (e21 = op2(e23, e25)) | ~ spl42_748), inference(backward_demodulation, [], [f4807, f5172])).
fof(f4807, plain, ~ (op2(e23, e25) = h6(e11)), inference(backward_demodulation, [], [f918, f1245])).
fof(f918, plain, ~ (op2(e23, e25) = op2(e25, e25)), inference(cnf_transformation, [], [f8])).
fof(f11277, plain, (~ spl42_597 | ~ spl42_748), inference(avatar_split_clause, [], [f11266, f5171, f4168])).
fof(f4168, plain, (spl42_597 <=> (e21 = op2(e21, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_597])])).
fof(f11266, plain, (~ (e21 = op2(e21, e25)) | ~ spl42_748), inference(backward_demodulation, [], [f4805, f5172])).
fof(f4805, plain, ~ (op2(e21, e25) = h6(e11)), inference(backward_demodulation, [], [f911, f1245])).
fof(f911, plain, ~ (op2(e21, e25) = op2(e25, e25)), inference(cnf_transformation, [], [f8])).
fof(f11263, plain, (~ spl42_504 | ~ spl42_1016), inference(avatar_split_clause, [], [f11256, f6684, f3782])).
fof(f11256, plain, (~ (e26 = op2(e23, e25)) | ~ spl42_1016), inference(backward_demodulation, [], [f4795, f6685])).
fof(f4795, plain, ~ (op2(e23, e25) = h6(e16)), inference(backward_demodulation, [], [f1025, f4785])).
fof(f1025, plain, ~ (op2(e23, e24) = op2(e23, e25)), inference(cnf_transformation, [], [f8])).
fof(f11261, plain, (~ spl42_539 | ~ spl42_1016), inference(avatar_split_clause, [], [f11254, f6684, f3927])).
fof(f3927, plain, (spl42_539 <=> (e26 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_539])])).
fof(f11254, plain, (~ (e26 = op2(e23, e20)) | ~ spl42_1016), inference(backward_demodulation, [], [f4792, f6685])).
fof(f4792, plain, ~ (op2(e23, e20) = h6(e16)), inference(backward_demodulation, [], [f1010, f4785])).
fof(f1010, plain, ~ (op2(e23, e20) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f11259, plain, (~ spl42_609 | ~ spl42_1016), inference(avatar_split_clause, [], [f11252, f6684, f4217])).
fof(f4217, plain, (spl42_609 <=> (e26 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_609])])).
fof(f11252, plain, (~ (e26 = op2(e21, e24)) | ~ spl42_1016), inference(backward_demodulation, [], [f4788, f6685])).
fof(f4788, plain, ~ (op2(e21, e24) = h6(e16)), inference(backward_demodulation, [], [f888, f4785])).
fof(f888, plain, ~ (op2(e21, e24) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f11244, plain, (~ spl42_670 | ~ spl42_656), inference(avatar_split_clause, [], [f11243, f4412, f4470])).
fof(f4470, plain, (spl42_670 <=> (e24 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_670])])).
fof(f11243, plain, (~ (e24 = op2(e20, e22)) | ~ spl42_656), inference(forward_demodulation, [], [f956, f4414])).
fof(f956, plain, ~ (op2(e20, e22) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f11236, plain, (~ spl42_663 | ~ spl42_656), inference(avatar_split_clause, [], [f11235, f4412, f4441])).
fof(f4441, plain, (spl42_663 <=> (e24 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_663])])).
fof(f11235, plain, (~ (e24 = op2(e20, e23)) | ~ spl42_656), inference(forward_demodulation, [], [f959, f4414])).
fof(f959, plain, ~ (op2(e20, e23) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f11225, plain, (spl42_644 | ~ spl42_392), inference(avatar_split_clause, [], [f9501, f3318, f4362])).
fof(f4362, plain, (spl42_644 <=> (e26 = op2(e20, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_644])])).
fof(f3318, plain, (spl42_392 <=> (e26 = op2(e26, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_392])])).
fof(f9501, plain, ((e26 = op2(e20, e26)) | ~ spl42_392), inference(backward_demodulation, [], [f4826, f9500])).
fof(f9500, plain, ((e20 = op2(h7(e11), e26)) | ~ spl42_392), inference(forward_demodulation, [], [f9491, f1252])).
fof(f9491, plain, ((e20 = op2(op2(e26, e26), e26)) | ~ spl42_392), inference(backward_demodulation, [], [f460, f3320])).
fof(f3320, plain, ((e26 = op2(e26, e20)) | ~ spl42_392), inference(avatar_component_clause, [], [f3318])).
fof(f11193, plain, (~ spl42_443 | ~ spl42_590), inference(avatar_contradiction_clause, [], [f11192])).
fof(f11192, plain, ($false | (~ spl42_443 | ~ spl42_590)), inference(subsumption_resolution, [], [f11191, f4141])).
fof(f4141, plain, ((e21 = op2(e21, e26)) | ~ spl42_590), inference(avatar_component_clause, [], [f4139])).
fof(f4139, plain, (spl42_590 <=> (e21 = op2(e21, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_590])])).
fof(f11191, plain, (~ (e21 = op2(e21, e26)) | ~ spl42_443), inference(forward_demodulation, [], [f931, f3532])).
fof(f931, plain, ~ (op2(e21, e26) = op2(e24, e26)), inference(cnf_transformation, [], [f8])).
fof(f11166, plain, (~ spl42_554 | ~ spl42_407), inference(avatar_split_clause, [], [f11165, f3381, f3990])).
fof(f3381, plain, (spl42_407 <=> (e20 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_407])])).
fof(f11165, plain, (~ (e20 = op2(e22, e24)) | ~ spl42_407), inference(forward_demodulation, [], [f894, f3383])).
fof(f3383, plain, ((e20 = op2(e25, e24)) | ~ spl42_407), inference(avatar_component_clause, [], [f3381])).
fof(f894, plain, ~ (op2(e22, e24) = op2(e25, e24)), inference(cnf_transformation, [], [f8])).
fof(f10944, plain, (~ spl42_320 | ~ spl42_313), inference(avatar_split_clause, [], [f10943, f2893, f2922])).
fof(f10943, plain, (~ (e14 = op1(e10, e13)) | ~ spl42_313), inference(forward_demodulation, [], [f665, f2895])).
fof(f665, plain, ~ (op1(e10, e13) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f10909, plain, (~ spl42_266 | ~ spl42_168), inference(avatar_split_clause, [], [f10908, f2292, f2698])).
fof(f2698, plain, (spl42_266 <=> (e16 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_266])])).
fof(f10908, plain, (~ (e16 = op1(e11, e14)) | ~ spl42_168), inference(forward_demodulation, [], [f594, f2294])).
fof(f594, plain, ~ (op1(e11, e14) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f10907, plain, (~ spl42_262 | ~ spl42_115), inference(avatar_split_clause, [], [f10906, f2073, f2682])).
fof(f2682, plain, (spl42_262 <=> (e12 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_262])])).
fof(f10906, plain, (~ (e12 = op1(e11, e14)) | ~ spl42_115), inference(forward_demodulation, [], [f595, f2075])).
fof(f595, plain, ~ (op1(e11, e14) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f10905, plain, (~ spl42_263 | ~ spl42_18), inference(avatar_split_clause, [], [f10904, f1671, f2686])).
fof(f10904, plain, (~ (e13 = op1(e11, e14)) | ~ spl42_18), inference(forward_demodulation, [], [f597, f1673])).
fof(f597, plain, ~ (op1(e11, e14) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f10858, plain, (~ spl42_217 | ~ spl42_168), inference(avatar_split_clause, [], [f10857, f2292, f2495])).
fof(f2495, plain, (spl42_217 <=> (e16 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_217])])).
fof(f10857, plain, (~ (e16 = op1(e12, e14)) | ~ spl42_168), inference(forward_demodulation, [], [f598, f2294])).
fof(f598, plain, ~ (op1(e12, e14) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f10850, plain, (spl42_211 | ~ spl42_115), inference(avatar_split_clause, [], [f10318, f2073, f2471])).
fof(f10318, plain, ((e10 = op1(e12, e14)) | ~ spl42_115), inference(backward_demodulation, [], [f4659, f2075])).
fof(f4659, plain, (e10 = op1(op1(e14, e14), e14)), inference(forward_demodulation, [], [f1196, f1200])).
fof(f1200, plain, (e14 = op1(e15, op1(e15, e15))), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((e16 = op1(op1(op1(op1(e15, op1(e15, e15)), op1(e15, op1(e15, e15))), e15), op1(e15, op1(e15, e15)))) & (e14 = op1(e15, op1(e15, e15))) & (e13 = op1(op1(op1(e15, op1(e15, e15)), op1(e15, op1(e15, e15))), e15)) & (e12 = op1(op1(e15, op1(e15, e15)), op1(e15, op1(e15, e15)))) & (e11 = op1(e15, e15)) & (e10 = op1(op1(op1(e15, op1(e15, e15)), op1(e15, op1(e15, e15))), op1(e15, op1(e15, e15))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax14)).
fof(f1196, plain, (e10 = op1(op1(op1(e15, op1(e15, e15)), op1(e15, op1(e15, e15))), op1(e15, op1(e15, e15)))), inference(cnf_transformation, [], [f14])).
fof(f10843, plain, (spl42_207 | ~ spl42_115), inference(avatar_split_clause, [], [f10317, f2073, f2454])).
fof(f10317, plain, ((e13 = op1(e12, e15)) | ~ spl42_115), inference(backward_demodulation, [], [f4653, f2075])).
fof(f4653, plain, (e13 = op1(op1(e14, e14), e15)), inference(forward_demodulation, [], [f1199, f1200])).
fof(f1199, plain, (e13 = op1(op1(op1(e15, op1(e15, e15)), op1(e15, op1(e15, e15))), e15)), inference(cnf_transformation, [], [f14])).
fof(f10804, plain, (~ spl42_182 | ~ spl42_168), inference(avatar_split_clause, [], [f10803, f2292, f2350])).
fof(f2350, plain, (spl42_182 <=> (e16 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_182])])).
fof(f10803, plain, (~ (e16 = op1(e13, e12)) | ~ spl42_168), inference(forward_demodulation, [], [f725, f2294])).
fof(f725, plain, ~ (op1(e13, e12) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f10800, plain, (~ spl42_156 | ~ spl42_58), inference(avatar_split_clause, [], [f10799, f1837, f2243])).
fof(f2243, plain, (spl42_156 <=> (e11 = op1(e13, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_156])])).
fof(f10799, plain, (~ (e11 = op1(e13, e15)) | ~ spl42_58), inference(forward_demodulation, [], [f624, f1839])).
fof(f624, plain, ~ (op1(e13, e15) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f10791, plain, (~ spl42_149 | ~ spl42_151), inference(avatar_contradiction_clause, [], [f10790])).
fof(f10790, plain, ($false | (~ spl42_149 | ~ spl42_151)), inference(subsumption_resolution, [], [f10789, f1098])).
fof(f10789, plain, ((e11 = e13) | (~ spl42_149 | ~ spl42_151)), inference(forward_demodulation, [], [f2224, f2216])).
fof(f2216, plain, ((e11 = op1(e13, e16)) | ~ spl42_149), inference(avatar_component_clause, [], [f2214])).
fof(f2214, plain, (spl42_149 <=> (e11 = op1(e13, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_149])])).
fof(f10780, plain, (~ spl42_136 | ~ spl42_115), inference(avatar_split_clause, [], [f10779, f2073, f2160])).
fof(f10779, plain, (~ (e12 = op1(e14, e11)) | ~ spl42_115), inference(forward_demodulation, [], [f742, f2075])).
fof(f742, plain, ~ (op1(e14, e11) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f10766, plain, (~ spl42_121 | ~ spl42_100), inference(avatar_split_clause, [], [f10765, f2011, f2098])).
fof(f2098, plain, (spl42_121 <=> (e11 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_121])])).
fof(f10765, plain, (~ (e11 = op1(e14, e13)) | ~ spl42_100), inference(forward_demodulation, [], [f751, f2013])).
fof(f751, plain, ~ (op1(e14, e13) = op1(e14, e16)), inference(cnf_transformation, [], [f7])).
fof(f10764, plain, (~ spl42_107 | ~ spl42_58), inference(avatar_split_clause, [], [f10763, f1837, f2040])).
fof(f10763, plain, (~ (e11 = op1(e14, e15)) | ~ spl42_58), inference(forward_demodulation, [], [f626, f1839])).
fof(f626, plain, ~ (op1(e14, e15) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f10737, plain, (spl42_89 | ~ spl42_58), inference(avatar_split_clause, [], [f10433, f1837, f1965])).
fof(f10433, plain, ((e14 = op1(e15, e11)) | ~ spl42_58), inference(backward_demodulation, [], [f1200, f1839])).
fof(f10734, plain, (~ spl42_79 | ~ spl42_58), inference(avatar_split_clause, [], [f10733, f1837, f1924])).
fof(f10733, plain, (~ (e11 = op1(e15, e12)) | ~ spl42_58), inference(forward_demodulation, [], [f768, f1839])).
fof(f768, plain, ~ (op1(e15, e12) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f10728, plain, (~ spl42_72 | ~ spl42_58), inference(avatar_split_clause, [], [f10727, f1837, f1895])).
fof(f1895, plain, (spl42_72 <=> (e11 = op1(e15, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_72])])).
fof(f10727, plain, (~ (e11 = op1(e15, e13)) | ~ spl42_58), inference(forward_demodulation, [], [f771, f1839])).
fof(f771, plain, ~ (op1(e15, e13) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f10722, plain, (~ spl42_65 | ~ spl42_58), inference(avatar_split_clause, [], [f10721, f1837, f1866])).
fof(f1866, plain, (spl42_65 <=> (e11 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_65])])).
fof(f10721, plain, (~ (e11 = op1(e15, e14)) | ~ spl42_58), inference(forward_demodulation, [], [f773, f1839])).
fof(f773, plain, ~ (op1(e15, e14) = op1(e15, e15)), inference(cnf_transformation, [], [f7])).
fof(f10646, plain, (spl42_271 | ~ spl42_23 | ~ spl42_168), inference(avatar_split_clause, [], [f10645, f2292, f1692, f2719])).
fof(f1692, plain, (spl42_23 <=> (e11 = op1(e16, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_23])])).
fof(f10645, plain, ((e14 = op1(e11, e13)) | (~ spl42_23 | ~ spl42_168)), inference(forward_demodulation, [], [f10201, f1694])).
fof(f1694, plain, ((e11 = op1(e16, e13)) | ~ spl42_23), inference(avatar_component_clause, [], [f1692])).
fof(f10561, plain, (spl42_26 | ~ spl42_28 | ~ spl42_168), inference(avatar_split_clause, [], [f10554, f2292, f1712, f1704])).
fof(f1712, plain, (spl42_28 <=> (e16 = op1(e16, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_28])])).
fof(f10554, plain, ((e14 = op1(e16, e13)) | (~ spl42_28 | ~ spl42_168)), inference(backward_demodulation, [], [f10201, f1714])).
fof(f1714, plain, ((e16 = op1(e16, e13)) | ~ spl42_28), inference(avatar_component_clause, [], [f1712])).
fof(f10443, plain, (~ spl42_58 | ~ spl42_85), inference(avatar_contradiction_clause, [], [f10442])).
fof(f10442, plain, ($false | (~ spl42_58 | ~ spl42_85)), inference(subsumption_resolution, [], [f10441, f1094])).
fof(f10441, plain, ((e10 = e14) | (~ spl42_58 | ~ spl42_85)), inference(forward_demodulation, [], [f10433, f1951])).
fof(f1951, plain, ((e10 = op1(e15, e11)) | ~ spl42_85), inference(avatar_component_clause, [], [f1949])).
fof(f1949, plain, (spl42_85 <=> (e10 = op1(e15, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_85])])).
fof(f10439, plain, (~ spl42_9 | ~ spl42_58), inference(avatar_split_clause, [], [f10431, f1837, f1634])).
fof(f10431, plain, (~ (e11 = op1(e16, e15)) | ~ spl42_58), inference(backward_demodulation, [], [f628, f1839])).
fof(f628, plain, ~ (op1(e15, e15) = op1(e16, e15)), inference(cnf_transformation, [], [f7])).
fof(f10366, plain, (~ spl42_71 | ~ spl42_92), inference(avatar_split_clause, [], [f10358, f1978, f1891])).
fof(f10358, plain, (~ (e10 = op1(e15, e13)) | ~ spl42_92), inference(backward_demodulation, [], [f757, f1980])).
fof(f757, plain, ~ (op1(e15, e10) = op1(e15, e13)), inference(cnf_transformation, [], [f7])).
fof(f10333, plain, (~ spl42_106 | ~ spl42_115 | ~ spl42_313), inference(avatar_contradiction_clause, [], [f10332])).
fof(f10332, plain, ($false | (~ spl42_106 | ~ spl42_115 | ~ spl42_313)), inference(subsumption_resolution, [], [f10331, f1104])).
fof(f10331, plain, ((e12 = e15) | (~ spl42_106 | ~ spl42_115 | ~ spl42_313)), inference(forward_demodulation, [], [f10330, f2075])).
fof(f10330, plain, ((e15 = op1(e14, e14)) | (~ spl42_106 | ~ spl42_313)), inference(forward_demodulation, [], [f10326, f2895])).
fof(f10326, plain, ((e15 = op1(op1(e10, e14), e14)) | ~ spl42_106), inference(backward_demodulation, [], [f297, f2038])).
fof(f2038, plain, ((e10 = op1(e14, e15)) | ~ spl42_106), inference(avatar_component_clause, [], [f2036])).
fof(f2036, plain, (spl42_106 <=> (e10 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_106])])).
fof(f10322, plain, (~ spl42_101 | ~ spl42_115), inference(avatar_split_clause, [], [f10315, f2073, f2015])).
fof(f2015, plain, (spl42_101 <=> (e12 = op1(e14, e16))), introduced(avatar_definition, [new_symbols(naming, [spl42_101])])).
fof(f10315, plain, (~ (e12 = op1(e14, e16)) | ~ spl42_115), inference(backward_demodulation, [], [f753, f2075])).
fof(f753, plain, ~ (op1(e14, e14) = op1(e14, e16)), inference(cnf_transformation, [], [f7])).
fof(f10321, plain, (~ spl42_108 | ~ spl42_115), inference(avatar_split_clause, [], [f10314, f2073, f2044])).
fof(f2044, plain, (spl42_108 <=> (e12 = op1(e14, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_108])])).
fof(f10314, plain, (~ (e12 = op1(e14, e15)) | ~ spl42_115), inference(backward_demodulation, [], [f752, f2075])).
fof(f752, plain, ~ (op1(e14, e14) = op1(e14, e15)), inference(cnf_transformation, [], [f7])).
fof(f10320, plain, (~ spl42_17 | ~ spl42_115), inference(avatar_split_clause, [], [f10313, f2073, f1667])).
fof(f1667, plain, (spl42_17 <=> (e12 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_17])])).
fof(f10313, plain, (~ (e12 = op1(e16, e14)) | ~ spl42_115), inference(backward_demodulation, [], [f606, f2075])).
fof(f606, plain, ~ (op1(e14, e14) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f10319, plain, (~ spl42_66 | ~ spl42_115), inference(avatar_split_clause, [], [f10312, f2073, f1870])).
fof(f1870, plain, (spl42_66 <=> (e12 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_66])])).
fof(f10312, plain, (~ (e12 = op1(e15, e14)) | ~ spl42_115), inference(backward_demodulation, [], [f605, f2075])).
fof(f605, plain, ~ (op1(e14, e14) = op1(e15, e14)), inference(cnf_transformation, [], [f7])).
fof(f10306, plain, (spl42_116 | ~ spl42_120 | ~ spl42_313), inference(avatar_split_clause, [], [f10305, f2893, f2094, f2077])).
fof(f2094, plain, (spl42_120 <=> (e10 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_120])])).
fof(f10305, plain, ((e13 = op1(e14, e14)) | (~ spl42_120 | ~ spl42_313)), inference(forward_demodulation, [], [f10299, f2895])).
fof(f10299, plain, ((e13 = op1(op1(e10, e14), e14)) | ~ spl42_120), inference(backward_demodulation, [], [f283, f2096])).
fof(f2096, plain, ((e10 = op1(e14, e13)) | ~ spl42_120), inference(avatar_component_clause, [], [f2094])).
fof(f10293, plain, (~ spl42_78 | ~ spl42_127), inference(avatar_split_clause, [], [f10285, f2123, f1920])).
fof(f1920, plain, (spl42_78 <=> (e10 = op1(e15, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_78])])).
fof(f10285, plain, (~ (e10 = op1(e15, e12)) | ~ spl42_127), inference(backward_demodulation, [], [f563, f2125])).
fof(f563, plain, ~ (op1(e14, e12) = op1(e15, e12)), inference(cnf_transformation, [], [f7])).
fof(f10276, plain, (spl42_114 | ~ spl42_134 | ~ spl42_313), inference(avatar_split_clause, [], [f10275, f2893, f2152, f2069])).
fof(f10275, plain, ((e11 = op1(e14, e14)) | (~ spl42_134 | ~ spl42_313)), inference(forward_demodulation, [], [f10267, f2895])).
fof(f10267, plain, ((e11 = op1(op1(e10, e14), e14)) | ~ spl42_134), inference(backward_demodulation, [], [f269, f2154])).
fof(f2154, plain, ((e10 = op1(e14, e11)) | ~ spl42_134), inference(avatar_component_clause, [], [f2152])).
fof(f269, plain, (e11 = op1(op1(op1(e14, e11), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f10255, plain, (spl42_113 | ~ spl42_141 | ~ spl42_313), inference(avatar_split_clause, [], [f10254, f2893, f2181, f2065])).
fof(f2065, plain, (spl42_113 <=> (e10 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_113])])).
fof(f2181, plain, (spl42_141 <=> (e10 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_141])])).
fof(f10254, plain, ((e10 = op1(e14, e14)) | (~ spl42_141 | ~ spl42_313)), inference(forward_demodulation, [], [f10244, f2895])).
fof(f10244, plain, ((e10 = op1(op1(e10, e14), e14)) | ~ spl42_141), inference(backward_demodulation, [], [f262, f2183])).
fof(f2183, plain, ((e10 = op1(e14, e10)) | ~ spl42_141), inference(avatar_component_clause, [], [f2181])).
fof(f262, plain, (e10 = op1(op1(op1(e14, e10), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f10220, plain, (~ spl42_155 | ~ spl42_316), inference(avatar_contradiction_clause, [], [f10219])).
fof(f10219, plain, ($false | (~ spl42_155 | ~ spl42_316)), inference(subsumption_resolution, [], [f10218, f1095])).
fof(f10218, plain, ((e10 = e15) | (~ spl42_155 | ~ spl42_316)), inference(forward_demodulation, [], [f10217, f2908])).
fof(f2908, plain, ((e10 = op1(e10, e13)) | ~ spl42_316), inference(avatar_component_clause, [], [f2906])).
fof(f2906, plain, (spl42_316 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_316])])).
fof(f10217, plain, ((e15 = op1(e10, e13)) | (~ spl42_155 | ~ spl42_316)), inference(forward_demodulation, [], [f10212, f2908])).
fof(f10212, plain, ((e15 = op1(op1(e10, e13), e13)) | ~ spl42_155), inference(backward_demodulation, [], [f296, f2241])).
fof(f10210, plain, (~ spl42_161 | ~ spl42_168), inference(avatar_split_clause, [], [f10205, f2292, f2263])).
fof(f10205, plain, (~ (e16 = op1(e13, e15)) | ~ spl42_168), inference(backward_demodulation, [], [f731, f2294])).
fof(f731, plain, ~ (op1(e13, e14) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f10209, plain, (~ spl42_21 | ~ spl42_168), inference(avatar_split_clause, [], [f10204, f2292, f1683])).
fof(f10204, plain, (~ (e16 = op1(e16, e14)) | ~ spl42_168), inference(backward_demodulation, [], [f604, f2294])).
fof(f604, plain, ~ (op1(e13, e14) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f10208, plain, (~ spl42_70 | ~ spl42_168), inference(avatar_split_clause, [], [f10203, f2292, f1886])).
fof(f1886, plain, (spl42_70 <=> (e16 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_70])])).
fof(f10203, plain, (~ (e16 = op1(e15, e14)) | ~ spl42_168), inference(backward_demodulation, [], [f603, f2294])).
fof(f603, plain, ~ (op1(e13, e14) = op1(e15, e14)), inference(cnf_transformation, [], [f7])).
fof(f10176, plain, (~ spl42_127 | ~ spl42_176), inference(avatar_split_clause, [], [f10165, f2326, f2123])).
fof(f2326, plain, (spl42_176 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl42_176])])).
fof(f10165, plain, (~ (e10 = op1(e14, e12)) | ~ spl42_176), inference(backward_demodulation, [], [f560, f2328])).
fof(f2328, plain, ((e10 = op1(e13, e12)) | ~ spl42_176), inference(avatar_component_clause, [], [f2326])).
fof(f560, plain, ~ (op1(e13, e12) = op1(e14, e12)), inference(cnf_transformation, [], [f7])).
fof(f10081, plain, (~ spl42_15 | ~ spl42_211), inference(avatar_split_clause, [], [f10071, f2471, f1659])).
fof(f1659, plain, (spl42_15 <=> (e10 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_15])])).
fof(f10071, plain, (~ (e10 = op1(e16, e14)) | ~ spl42_211), inference(backward_demodulation, [], [f601, f2473])).
fof(f601, plain, ~ (op1(e12, e14) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f10080, plain, (~ spl42_64 | ~ spl42_211), inference(avatar_split_clause, [], [f10070, f2471, f1862])).
fof(f1862, plain, (spl42_64 <=> (e10 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_64])])).
fof(f10070, plain, (~ (e10 = op1(e15, e14)) | ~ spl42_211), inference(backward_demodulation, [], [f600, f2473])).
fof(f600, plain, ~ (op1(e12, e14) = op1(e15, e14)), inference(cnf_transformation, [], [f7])).
fof(f10079, plain, (~ spl42_113 | ~ spl42_211), inference(avatar_split_clause, [], [f10069, f2471, f2065])).
fof(f10069, plain, (~ (e10 = op1(e14, e14)) | ~ spl42_211), inference(backward_demodulation, [], [f599, f2473])).
fof(f599, plain, ~ (op1(e12, e14) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f10077, plain, (~ spl42_211 | ~ spl42_323), inference(avatar_contradiction_clause, [], [f10076])).
fof(f10076, plain, ($false | (~ spl42_211 | ~ spl42_323)), inference(subsumption_resolution, [], [f10075, f1094])).
fof(f10075, plain, ((e10 = e14) | (~ spl42_211 | ~ spl42_323)), inference(forward_demodulation, [], [f10074, f2937])).
fof(f2937, plain, ((e10 = op1(e10, e12)) | ~ spl42_323), inference(avatar_component_clause, [], [f2935])).
fof(f10074, plain, ((e14 = op1(e10, e12)) | (~ spl42_211 | ~ spl42_323)), inference(forward_demodulation, [], [f10067, f2937])).
fof(f10067, plain, ((e14 = op1(op1(e10, e12), e12)) | ~ spl42_211), inference(backward_demodulation, [], [f288, f2473])).
fof(f10064, plain, (~ spl42_211 | ~ spl42_218), inference(avatar_split_clause, [], [f10053, f2500, f2471])).
fof(f2500, plain, (spl42_218 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_218])])).
fof(f10053, plain, (~ (e10 = op1(e12, e14)) | ~ spl42_218), inference(backward_demodulation, [], [f707, f2502])).
fof(f2502, plain, ((e10 = op1(e12, e13)) | ~ spl42_218), inference(avatar_component_clause, [], [f2500])).
fof(f707, plain, ~ (op1(e12, e13) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f10023, plain, (~ spl42_211 | ~ spl42_232), inference(avatar_split_clause, [], [f10010, f2558, f2471])).
fof(f2558, plain, (spl42_232 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_232])])).
fof(f10010, plain, (~ (e10 = op1(e12, e14)) | ~ spl42_232), inference(backward_demodulation, [], [f700, f2560])).
fof(f2560, plain, ((e10 = op1(e12, e11)) | ~ spl42_232), inference(avatar_component_clause, [], [f2558])).
fof(f700, plain, ~ (op1(e12, e11) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f10000, plain, (~ spl42_211 | ~ spl42_239), inference(avatar_split_clause, [], [f9989, f2587, f2471])).
fof(f2587, plain, (spl42_239 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_239])])).
fof(f9989, plain, (~ (e10 = op1(e12, e14)) | ~ spl42_239), inference(backward_demodulation, [], [f695, f2589])).
fof(f2589, plain, ((e10 = op1(e12, e10)) | ~ spl42_239), inference(avatar_component_clause, [], [f2587])).
fof(f695, plain, ~ (op1(e12, e10) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f9931, plain, (~ spl42_211 | ~ spl42_260), inference(avatar_split_clause, [], [f9920, f2674, f2471])).
fof(f2674, plain, (spl42_260 <=> (e10 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_260])])).
fof(f9920, plain, (~ (e10 = op1(e12, e14)) | ~ spl42_260), inference(backward_demodulation, [], [f593, f2676])).
fof(f2676, plain, ((e10 = op1(e11, e14)) | ~ spl42_260), inference(avatar_component_clause, [], [f2674])).
fof(f593, plain, ~ (op1(e11, e14) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f9793, plain, (~ spl42_155 | ~ spl42_302), inference(avatar_split_clause, [], [f9783, f2848, f2239])).
fof(f2848, plain, (spl42_302 <=> (e10 = op1(e10, e15))), introduced(avatar_definition, [new_symbols(naming, [spl42_302])])).
fof(f9783, plain, (~ (e10 = op1(e13, e15)) | ~ spl42_302), inference(backward_demodulation, [], [f610, f2850])).
fof(f2850, plain, ((e10 = op1(e10, e15)) | ~ spl42_302), inference(avatar_component_clause, [], [f2848])).
fof(f610, plain, ~ (op1(e10, e15) = op1(e13, e15)), inference(cnf_transformation, [], [f7])).
fof(f9777, plain, (~ spl42_19 | ~ spl42_313), inference(avatar_split_clause, [], [f9769, f2893, f1675])).
fof(f1675, plain, (spl42_19 <=> (e14 = op1(e16, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_19])])).
fof(f9769, plain, (~ (e14 = op1(e16, e14)) | ~ spl42_313), inference(backward_demodulation, [], [f592, f2895])).
fof(f592, plain, ~ (op1(e10, e14) = op1(e16, e14)), inference(cnf_transformation, [], [f7])).
fof(f9776, plain, (~ spl42_68 | ~ spl42_313), inference(avatar_split_clause, [], [f9768, f2893, f1878])).
fof(f1878, plain, (spl42_68 <=> (e14 = op1(e15, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_68])])).
fof(f9768, plain, (~ (e14 = op1(e15, e14)) | ~ spl42_313), inference(backward_demodulation, [], [f591, f2895])).
fof(f591, plain, ~ (op1(e10, e14) = op1(e15, e14)), inference(cnf_transformation, [], [f7])).
fof(f9775, plain, (~ spl42_117 | ~ spl42_313), inference(avatar_split_clause, [], [f9767, f2893, f2081])).
fof(f2081, plain, (spl42_117 <=> (e14 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_117])])).
fof(f9767, plain, (~ (e14 = op1(e14, e14)) | ~ spl42_313), inference(backward_demodulation, [], [f590, f2895])).
fof(f590, plain, ~ (op1(e10, e14) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f9773, plain, (~ spl42_215 | ~ spl42_313), inference(avatar_split_clause, [], [f9765, f2893, f2487])).
fof(f2487, plain, (spl42_215 <=> (e14 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_215])])).
fof(f9765, plain, (~ (e14 = op1(e12, e14)) | ~ spl42_313), inference(backward_demodulation, [], [f588, f2895])).
fof(f588, plain, ~ (op1(e10, e14) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f9772, plain, (~ spl42_264 | ~ spl42_313), inference(avatar_split_clause, [], [f9764, f2893, f2690])).
fof(f2690, plain, (spl42_264 <=> (e14 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl42_264])])).
fof(f9764, plain, (~ (e14 = op1(e11, e14)) | ~ spl42_313), inference(backward_demodulation, [], [f587, f2895])).
fof(f587, plain, ~ (op1(e10, e14) = op1(e11, e14)), inference(cnf_transformation, [], [f7])).
fof(f9578, plain, (~ spl42_376 | ~ spl42_377), inference(avatar_contradiction_clause, [], [f9577])).
fof(f9577, plain, ($false | (~ spl42_376 | ~ spl42_377)), inference(subsumption_resolution, [], [f9576, f1130])).
fof(f1130, plain, ~ (e24 = e25), inference(cnf_transformation, [], [f10])).
fof(f9576, plain, ((e24 = e25) | (~ spl42_376 | ~ spl42_377)), inference(backward_demodulation, [], [f3258, f3254])).
fof(f3258, plain, ((e25 = op2(e26, e22)) | ~ spl42_377), inference(avatar_component_clause, [], [f3256])).
fof(f3256, plain, (spl42_377 <=> (e25 = op2(e26, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_377])])).
fof(f9473, plain, (spl42_748 | ~ spl42_401), inference(avatar_split_clause, [], [f9472, f3356, f5171])).
fof(f9472, plain, ((e21 = h6(e11)) | ~ spl42_401), inference(backward_demodulation, [], [f1245, f3358])).
fof(f9359, plain, (~ spl42_449 | ~ spl42_458 | ~ spl42_656), inference(avatar_contradiction_clause, [], [f9358])).
fof(f9358, plain, ($false | (~ spl42_449 | ~ spl42_458 | ~ spl42_656)), inference(subsumption_resolution, [], [f9357, f1125])).
fof(f9357, plain, ((e22 = e25) | (~ spl42_449 | ~ spl42_458 | ~ spl42_656)), inference(forward_demodulation, [], [f9356, f3594])).
fof(f9356, plain, ((e25 = op2(e24, e24)) | (~ spl42_449 | ~ spl42_656)), inference(forward_demodulation, [], [f9351, f4414])).
fof(f9351, plain, ((e25 = op2(op2(e20, e24), e24)) | ~ spl42_449), inference(backward_demodulation, [], [f493, f3557])).
fof(f3557, plain, ((e20 = op2(e24, e25)) | ~ spl42_449), inference(avatar_component_clause, [], [f3555])).
fof(f3555, plain, (spl42_449 <=> (e20 = op2(e24, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_449])])).
fof(f9350, plain, (spl42_743 | ~ spl42_458), inference(avatar_split_clause, [], [f9349, f3592, f5142])).
fof(f9349, plain, ((e22 = h5(e11)) | ~ spl42_458), inference(backward_demodulation, [], [f1238, f3594])).
fof(f9326, plain, (~ spl42_421 | ~ spl42_470), inference(avatar_split_clause, [], [f9317, f3642, f3439])).
fof(f3439, plain, (spl42_421 <=> (e20 = op2(e25, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_421])])).
fof(f9317, plain, (~ (e20 = op2(e25, e22)) | ~ spl42_470), inference(backward_demodulation, [], [f857, f3644])).
fof(f857, plain, ~ (op2(e24, e22) = op2(e25, e22)), inference(cnf_transformation, [], [f8])).
fof(f9307, plain, (spl42_457 | ~ spl42_477 | ~ spl42_656), inference(avatar_split_clause, [], [f9306, f4412, f3671, f3588])).
fof(f3671, plain, (spl42_477 <=> (e20 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_477])])).
fof(f9306, plain, ((e21 = op2(e24, e24)) | (~ spl42_477 | ~ spl42_656)), inference(forward_demodulation, [], [f9297, f4414])).
fof(f9297, plain, ((e21 = op2(op2(e20, e24), e24)) | ~ spl42_477), inference(backward_demodulation, [], [f465, f3673])).
fof(f3673, plain, ((e20 = op2(e24, e21)) | ~ spl42_477), inference(avatar_component_clause, [], [f3671])).
fof(f465, plain, (e21 = op2(op2(op2(e24, e21), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f9284, plain, (spl42_456 | ~ spl42_484 | ~ spl42_656), inference(avatar_split_clause, [], [f9283, f4412, f3700, f3584])).
fof(f3700, plain, (spl42_484 <=> (e20 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_484])])).
fof(f9283, plain, ((e20 = op2(e24, e24)) | (~ spl42_484 | ~ spl42_656)), inference(forward_demodulation, [], [f9272, f4414])).
fof(f9272, plain, ((e20 = op2(op2(e20, e24), e24)) | ~ spl42_484), inference(backward_demodulation, [], [f458, f3702])).
fof(f3702, plain, ((e20 = op2(e24, e20)) | ~ spl42_484), inference(avatar_component_clause, [], [f3700])).
fof(f9242, plain, (~ spl42_498 | ~ spl42_659), inference(avatar_contradiction_clause, [], [f9241])).
fof(f9241, plain, ($false | (~ spl42_498 | ~ spl42_659)), inference(subsumption_resolution, [], [f9240, f1116])).
fof(f9240, plain, ((e20 = e25) | (~ spl42_498 | ~ spl42_659)), inference(forward_demodulation, [], [f9239, f4427])).
fof(f4427, plain, ((e20 = op2(e20, e23)) | ~ spl42_659), inference(avatar_component_clause, [], [f4425])).
fof(f4425, plain, (spl42_659 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_659])])).
fof(f9239, plain, ((e25 = op2(e20, e23)) | (~ spl42_498 | ~ spl42_659)), inference(forward_demodulation, [], [f9232, f4427])).
fof(f9232, plain, ((e25 = op2(op2(e20, e23), e23)) | ~ spl42_498), inference(backward_demodulation, [], [f492, f3760])).
fof(f492, plain, (e25 = op2(op2(op2(e23, e25), e23), e23)), inference(cnf_transformation, [], [f6])).
fof(f9231, plain, (spl42_1016 | ~ spl42_511), inference(avatar_split_clause, [], [f9228, f3811, f6684])).
fof(f9228, plain, ((e26 = h6(e16)) | ~ spl42_511), inference(backward_demodulation, [], [f4785, f3813])).
fof(f9174, plain, (~ spl42_498 | ~ spl42_533), inference(avatar_split_clause, [], [f9163, f3903, f3758])).
fof(f3903, plain, (spl42_533 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_533])])).
fof(f9163, plain, (~ (e20 = op2(e23, e25)) | ~ spl42_533), inference(backward_demodulation, [], [f1011, f3905])).
fof(f3905, plain, ((e20 = op2(e23, e20)) | ~ spl42_533), inference(avatar_component_clause, [], [f3903])).
fof(f1011, plain, ~ (op2(e23, e20) = op2(e23, e25)), inference(cnf_transformation, [], [f8])).
fof(f9147, plain, (~ spl42_544 | ~ spl42_545), inference(avatar_contradiction_clause, [], [f9146])).
fof(f9146, plain, ($false | (~ spl42_544 | ~ spl42_545)), inference(subsumption_resolution, [], [f9145, f1130])).
fof(f9145, plain, ((e24 = e25) | (~ spl42_544 | ~ spl42_545)), inference(backward_demodulation, [], [f3954, f3950])).
fof(f3954, plain, ((e25 = op2(e22, e26)) | ~ spl42_545), inference(avatar_component_clause, [], [f3952])).
fof(f3952, plain, (spl42_545 <=> (e25 = op2(e22, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_545])])).
fof(f9111, plain, (~ spl42_871 | ~ spl42_554), inference(avatar_split_clause, [], [f9100, f3990, f5801])).
fof(f5801, plain, (spl42_871 <=> (e20 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl42_871])])).
fof(f9100, plain, (~ (e20 = h3(e11)) | ~ spl42_554), inference(backward_demodulation, [], [f4731, f3992])).
fof(f4731, plain, ~ (op2(e22, e24) = h3(e11)), inference(backward_demodulation, [], [f998, f1224])).
fof(f998, plain, ~ (op2(e22, e22) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f9110, plain, (~ spl42_540 | ~ spl42_554), inference(avatar_split_clause, [], [f9099, f3990, f3932])).
fof(f3932, plain, (spl42_540 <=> (e20 = op2(e22, e26))), introduced(avatar_definition, [new_symbols(naming, [spl42_540])])).
fof(f9099, plain, (~ (e20 = op2(e22, e26)) | ~ spl42_554), inference(backward_demodulation, [], [f1005, f3992])).
fof(f1005, plain, ~ (op2(e22, e24) = op2(e22, e26)), inference(cnf_transformation, [], [f8])).
fof(f9108, plain, (~ spl42_358 | ~ spl42_554), inference(avatar_split_clause, [], [f9097, f3990, f3178])).
fof(f9097, plain, (~ (e20 = op2(e26, e24)) | ~ spl42_554), inference(backward_demodulation, [], [f895, f3992])).
fof(f895, plain, ~ (op2(e22, e24) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f9106, plain, (~ spl42_554 | ~ spl42_666), inference(avatar_contradiction_clause, [], [f9105])).
fof(f9105, plain, ($false | (~ spl42_554 | ~ spl42_666)), inference(subsumption_resolution, [], [f9104, f1115])).
fof(f9104, plain, ((e20 = e24) | (~ spl42_554 | ~ spl42_666)), inference(forward_demodulation, [], [f9103, f4456])).
fof(f4456, plain, ((e20 = op2(e20, e22)) | ~ spl42_666), inference(avatar_component_clause, [], [f4454])).
fof(f9103, plain, ((e24 = op2(e20, e22)) | (~ spl42_554 | ~ spl42_666)), inference(forward_demodulation, [], [f9095, f4456])).
fof(f9095, plain, ((e24 = op2(op2(e20, e22), e22)) | ~ spl42_554), inference(backward_demodulation, [], [f484, f3992])).
fof(f9090, plain, (~ spl42_554 | ~ spl42_561), inference(avatar_split_clause, [], [f9078, f4019, f3990])).
fof(f4019, plain, (spl42_561 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_561])])).
fof(f9078, plain, (~ (e20 = op2(e22, e24)) | ~ spl42_561), inference(backward_demodulation, [], [f1001, f4021])).
fof(f4021, plain, ((e20 = op2(e22, e23)) | ~ spl42_561), inference(avatar_component_clause, [], [f4019])).
fof(f1001, plain, ~ (op2(e22, e23) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f9073, plain, (spl42_871 | ~ spl42_568), inference(avatar_split_clause, [], [f9072, f4048, f5801])).
fof(f9072, plain, ((e20 = h3(e11)) | ~ spl42_568), inference(backward_demodulation, [], [f1224, f4050])).
fof(f4050, plain, ((e20 = op2(e22, e22)) | ~ spl42_568), inference(avatar_component_clause, [], [f4048])).
fof(f9066, plain, (~ spl42_554 | ~ spl42_575), inference(avatar_split_clause, [], [f9052, f4077, f3990])).
fof(f4077, plain, (spl42_575 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_575])])).
fof(f9052, plain, (~ (e20 = op2(e22, e24)) | ~ spl42_575), inference(backward_demodulation, [], [f994, f4079])).
fof(f4079, plain, ((e20 = op2(e22, e21)) | ~ spl42_575), inference(avatar_component_clause, [], [f4077])).
fof(f994, plain, ~ (op2(e22, e21) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f9041, plain, (~ spl42_554 | ~ spl42_582), inference(avatar_split_clause, [], [f9029, f4106, f3990])).
fof(f4106, plain, (spl42_582 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_582])])).
fof(f9029, plain, (~ (e20 = op2(e22, e24)) | ~ spl42_582), inference(backward_demodulation, [], [f989, f4108])).
fof(f4108, plain, ((e20 = op2(e22, e20)) | ~ spl42_582), inference(avatar_component_clause, [], [f4106])).
fof(f989, plain, ~ (op2(e22, e20) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f8967, plain, (~ spl42_554 | ~ spl42_603), inference(avatar_split_clause, [], [f8955, f4193, f3990])).
fof(f4193, plain, (spl42_603 <=> (e20 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_603])])).
fof(f8955, plain, (~ (e20 = op2(e22, e24)) | ~ spl42_603), inference(backward_demodulation, [], [f887, f4195])).
fof(f4195, plain, ((e20 = op2(e21, e24)) | ~ spl42_603), inference(avatar_component_clause, [], [f4193])).
fof(f887, plain, ~ (op2(e21, e24) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f8865, plain, (~ spl42_399 | ~ spl42_644), inference(avatar_split_clause, [], [f8858, f4362, f3347])).
fof(f8858, plain, (~ (e26 = op2(e25, e26)) | ~ spl42_644), inference(backward_demodulation, [], [f927, f4364])).
fof(f4364, plain, ((e26 = op2(e20, e26)) | ~ spl42_644), inference(avatar_component_clause, [], [f4362])).
fof(f927, plain, ~ (op2(e20, e26) = op2(e25, e26)), inference(cnf_transformation, [], [f8])).
fof(f8847, plain, (~ spl42_498 | ~ spl42_645), inference(avatar_split_clause, [], [f8837, f4367, f3758])).
fof(f4367, plain, (spl42_645 <=> (e20 = op2(e20, e25))), introduced(avatar_definition, [new_symbols(naming, [spl42_645])])).
fof(f8837, plain, (~ (e20 = op2(e23, e25)) | ~ spl42_645), inference(backward_demodulation, [], [f904, f4369])).
fof(f4369, plain, ((e20 = op2(e20, e25)) | ~ spl42_645), inference(avatar_component_clause, [], [f4367])).
fof(f904, plain, ~ (op2(e20, e25) = op2(e23, e25)), inference(cnf_transformation, [], [f8])).
fof(f8831, plain, (~ spl42_927 | ~ spl42_656), inference(avatar_split_clause, [], [f8822, f4412, f6082])).
fof(f8822, plain, (~ (e24 = h1(e11)) | ~ spl42_656), inference(backward_demodulation, [], [f4687, f4414])).
fof(f4687, plain, ~ (op2(e20, e24) = h1(e11)), inference(backward_demodulation, [], [f947, f1210])).
fof(f947, plain, ~ (op2(e20, e20) = op2(e20, e24)), inference(cnf_transformation, [], [f8])).
fof(f8828, plain, (~ spl42_362 | ~ spl42_656), inference(avatar_split_clause, [], [f8819, f4412, f3194])).
fof(f3194, plain, (spl42_362 <=> (e24 = op2(e26, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_362])])).
fof(f8819, plain, (~ (e24 = op2(e26, e24)) | ~ spl42_656), inference(backward_demodulation, [], [f886, f4414])).
fof(f886, plain, ~ (op2(e20, e24) = op2(e26, e24)), inference(cnf_transformation, [], [f8])).
fof(f8827, plain, (~ spl42_411 | ~ spl42_656), inference(avatar_split_clause, [], [f8818, f4412, f3397])).
fof(f3397, plain, (spl42_411 <=> (e24 = op2(e25, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_411])])).
fof(f8818, plain, (~ (e24 = op2(e25, e24)) | ~ spl42_656), inference(backward_demodulation, [], [f885, f4414])).
fof(f885, plain, ~ (op2(e20, e24) = op2(e25, e24)), inference(cnf_transformation, [], [f8])).
fof(f8826, plain, (~ spl42_558 | ~ spl42_656), inference(avatar_split_clause, [], [f8817, f4412, f4006])).
fof(f4006, plain, (spl42_558 <=> (e24 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_558])])).
fof(f8817, plain, (~ (e24 = op2(e22, e24)) | ~ spl42_656), inference(backward_demodulation, [], [f882, f4414])).
fof(f882, plain, ~ (op2(e20, e24) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f8825, plain, (~ spl42_607 | ~ spl42_656), inference(avatar_split_clause, [], [f8816, f4412, f4209])).
fof(f4209, plain, (spl42_607 <=> (e24 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_607])])).
fof(f8816, plain, (~ (e24 = op2(e21, e24)) | ~ spl42_656), inference(backward_demodulation, [], [f881, f4414])).
fof(f881, plain, ~ (op2(e20, e24) = op2(e21, e24)), inference(cnf_transformation, [], [f8])).
fof(f6867, plain, (~ spl42_1012 | ~ spl42_1013 | ~ spl42_1014 | ~ spl42_1015 | spl42_749 | spl42_745 | spl42_741 | spl42_737 | spl42_733 | spl42_729 | ~ spl42_1016 | ~ spl42_1017 | ~ spl42_1018 | ~ spl42_1019 | ~ spl42_1020 | ~ spl42_1021 | ~ spl42_1022 | ~ spl42_1023 | ~ spl42_1024 | ~ spl42_1025 | ~ spl42_1026 | ~ spl42_1027 | ~ spl42_1028 | ~ spl42_1029 | ~ spl42_1030 | ~ spl42_1031 | ~ spl42_1032 | ~ spl42_1033 | ~ spl42_1034 | ~ spl42_1035 | ~ spl42_1036 | ~ spl42_1037 | ~ spl42_1038 | ~ spl42_1039 | ~ spl42_1040 | ~ spl42_1041 | ~ spl42_1042 | ~ spl42_1043 | ~ spl42_1044 | ~ spl42_1045 | ~ spl42_1046 | ~ spl42_1047 | ~ spl42_1048 | ~ spl42_1049 | ~ spl42_1050 | ~ spl42_1051 | ~ spl42_1052 | ~ spl42_1053 | ~ spl42_1054 | ~ spl42_1055 | ~ spl42_1056 | ~ spl42_1057 | ~ spl42_1058 | ~ spl42_1059 | ~ spl42_1060 | ~ spl42_1061), inference(avatar_split_clause, [], [f6666, f6864, f6860, f6856, f6852, f6848, f6844, f6840, f6836, f6832, f6828, f6824, f6820, f6816, f6812, f6808, f6804, f6800, f6796, f6792, f6788, f6784, f6780, f6776, f6772, f6768, f6764, f6760, f6756, f6752, f6748, f6744, f6740, f6736, f6732, f6728, f6724, f6720, f6716, f6712, f6708, f6704, f6700, f6696, f6692, f6688, f6684, f5054, f5079, f5104, f5129, f5153, f5177, f6680, f6676, f6672, f6668])).
fof(f5177, plain, (spl42_749 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl42_749])])).
fof(f5153, plain, (spl42_745 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl42_745])])).
fof(f5129, plain, (spl42_741 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl42_741])])).
fof(f5104, plain, (spl42_737 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl42_737])])).
fof(f5079, plain, (spl42_733 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl42_733])])).
fof(f5054, plain, (spl42_729 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl42_729])])).
fof(f6666, plain, (~ (h1(e11) = h6(op1(e10, e10))) | ~ (h6(op1(e10, e11)) = op2(e20, h6(e11))) | ~ (h6(op1(e10, e12)) = op2(e20, h5(e11))) | ~ (op2(e20, e23) = h6(op1(e10, e13))) | ~ (op2(e20, e24) = h6(op1(e10, e14))) | ~ (op2(e20, e25) = h6(op1(e10, e15))) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11)))), inference(forward_demodulation, [], [f6665, f1210])).
fof(f6665, plain, (~ (op2(e20, e20) = h6(op1(e10, e10))) | ~ (h6(op1(e10, e11)) = op2(e20, h6(e11))) | ~ (h6(op1(e10, e12)) = op2(e20, h5(e11))) | ~ (op2(e20, e23) = h6(op1(e10, e13))) | ~ (op2(e20, e24) = h6(op1(e10, e14))) | ~ (op2(e20, e25) = h6(op1(e10, e15))) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11)))), inference(forward_demodulation, [], [f6664, f4820])).
fof(f6664, plain, (~ (h6(op1(e10, e11)) = op2(e20, h6(e11))) | ~ (h6(op1(e10, e12)) = op2(e20, h5(e11))) | ~ (op2(e20, e23) = h6(op1(e10, e13))) | ~ (op2(e20, e24) = h6(op1(e10, e14))) | ~ (op2(e20, e25) = h6(op1(e10, e15))) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6663, f4820])).
fof(f6663, plain, (~ (h6(op1(e10, e12)) = op2(e20, h5(e11))) | ~ (op2(e20, e23) = h6(op1(e10, e13))) | ~ (op2(e20, e24) = h6(op1(e10, e14))) | ~ (op2(e20, e25) = h6(op1(e10, e15))) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6662, f4820])).
fof(f6662, plain, (~ (h6(op1(e10, e12)) = op2(h6(e10), h5(e11))) | ~ (op2(e20, e23) = h6(op1(e10, e13))) | ~ (op2(e20, e24) = h6(op1(e10, e14))) | ~ (op2(e20, e25) = h6(op1(e10, e15))) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6661, f4802])).
fof(f6661, plain, (~ (op2(e20, e23) = h6(op1(e10, e13))) | ~ (op2(e20, e24) = h6(op1(e10, e14))) | ~ (op2(e20, e25) = h6(op1(e10, e15))) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6660, f4820])).
fof(f6660, plain, (~ (h6(op1(e10, e13)) = op2(h6(e10), e23)) | ~ (op2(e20, e24) = h6(op1(e10, e14))) | ~ (op2(e20, e25) = h6(op1(e10, e15))) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6659, f4800])).
fof(f6659, plain, (~ (op2(e20, e24) = h6(op1(e10, e14))) | ~ (op2(e20, e25) = h6(op1(e10, e15))) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6658, f4820])).
fof(f6658, plain, (~ (h6(op1(e10, e14)) = op2(h6(e10), e24)) | ~ (op2(e20, e25) = h6(op1(e10, e15))) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6657, f4797])).
fof(f6657, plain, (~ (op2(e20, e25) = h6(op1(e10, e15))) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6656, f4820])).
fof(f6656, plain, (~ (h6(op1(e10, e15)) = op2(h6(e10), e25)) | ~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6655, f1243])).
fof(f6655, plain, (~ (h6(op1(e10, e16)) = op2(e20, h6(e16))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6654, f4820])).
fof(f6654, plain, (~ (h6(op1(e11, e10)) = op2(h6(e11), e20)) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6653, f4820])).
fof(f6653, plain, (~ (h6(op1(e11, e12)) = op2(h6(e11), h5(e11))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6652, f4802])).
fof(f6652, plain, (~ (h6(op1(e11, e13)) = op2(h6(e11), e23)) | ~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6651, f4800])).
fof(f6651, plain, (~ (h6(op1(e11, e14)) = op2(h6(e11), e24)) | ~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6650, f4797])).
fof(f6650, plain, (~ (h6(op1(e11, e15)) = op2(h6(e11), e25)) | ~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6649, f1243])).
fof(f6649, plain, (~ (h6(op1(e12, e10)) = op2(h5(e11), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6648, f4802])).
fof(f6648, plain, (~ (h6(op1(e12, e10)) = op2(h6(e12), e20)) | ~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6647, f4820])).
fof(f6647, plain, (~ (h6(op1(e12, e11)) = op2(h5(e11), h6(e11))) | ~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6646, f4802])).
fof(f6646, plain, (~ (op2(h5(e11), h5(e11)) = h6(op1(e12, e12))) | ~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6645, f4802])).
fof(f6645, plain, (~ (h6(op1(e12, e13)) = op2(h5(e11), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6644, f4802])).
fof(f6644, plain, (~ (h6(op1(e12, e13)) = op2(h6(e12), e23)) | ~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6643, f4800])).
fof(f6643, plain, (~ (e20 = h6(op1(e12, e14))) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6642, f4779])).
fof(f6642, plain, (~ (h6(op1(e12, e14)) = op2(h5(e11), e24)) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6641, f4802])).
fof(f6641, plain, (~ (h6(op1(e12, e14)) = op2(h6(e12), e24)) | ~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6640, f4797])).
fof(f6640, plain, (~ (e23 = h6(op1(e12, e15))) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6639, f4778])).
fof(f6639, plain, (~ (h6(op1(e12, e15)) = op2(h5(e11), e25)) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6638, f4802])).
fof(f6638, plain, (~ (h6(op1(e12, e15)) = op2(h6(e12), e25)) | ~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6637, f1243])).
fof(f6637, plain, (~ (h6(op1(e12, e16)) = op2(h5(e11), h6(e16))) | ~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6636, f4802])).
fof(f6636, plain, (~ (op2(e23, e20) = h6(op1(e13, e10))) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6635, f4800])).
fof(f6635, plain, (~ (h6(op1(e13, e10)) = op2(h6(e13), e20)) | ~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6634, f4820])).
fof(f6634, plain, (~ (h6(op1(e13, e11)) = op2(e23, h6(e11))) | ~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6633, f4800])).
fof(f6633, plain, (~ (h6(op1(e13, e12)) = op2(e23, h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6632, f4800])).
fof(f6632, plain, (~ (h6(op1(e13, e12)) = op2(h6(e13), h5(e11))) | ~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6631, f4802])).
fof(f6631, plain, (~ (h4(e11) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6630, f1231])).
fof(f6630, plain, (~ (op2(e23, e23) = h6(op1(e13, e13))) | ~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6629, f4800])).
fof(f6629, plain, (~ (h6(e16) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6628, f4785])).
fof(f6628, plain, (~ (op2(e23, e24) = h6(op1(e13, e14))) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6627, f4800])).
fof(f6627, plain, (~ (h6(op1(e13, e14)) = op2(h6(e13), e24)) | ~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6626, f4797])).
fof(f6626, plain, (~ (op2(e23, e25) = h6(op1(e13, e15))) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6625, f4800])).
fof(f6625, plain, (~ (h6(op1(e13, e15)) = op2(h6(e13), e25)) | ~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6624, f1243])).
fof(f6624, plain, (~ (h6(op1(e13, e16)) = op2(e23, h6(e16))) | ~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6623, f4800])).
fof(f6623, plain, (~ (op2(e24, e20) = h6(op1(e14, e10))) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6622, f4797])).
fof(f6622, plain, (~ (h6(op1(e14, e10)) = op2(h6(e14), e20)) | ~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6621, f4820])).
fof(f6621, plain, (~ (h6(op1(e14, e11)) = op2(e24, h6(e11))) | ~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6620, f4797])).
fof(f6620, plain, (~ (h5(e14) = h6(op1(e14, e12))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6619, f4777])).
fof(f6619, plain, (~ (h6(op1(e14, e12)) = op2(e24, h5(e11))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6618, f4797])).
fof(f6618, plain, (~ (h6(op1(e14, e12)) = op2(h6(e14), h5(e11))) | ~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6617, f4802])).
fof(f6617, plain, (~ (op2(e24, e23) = h6(op1(e14, e13))) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6616, f4797])).
fof(f6616, plain, (~ (h6(op1(e14, e13)) = op2(h6(e14), e23)) | ~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6615, f4800])).
fof(f6615, plain, (~ (h5(e11) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6614, f1238])).
fof(f6614, plain, (~ (op2(e24, e24) = h6(op1(e14, e14))) | ~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6613, f4797])).
fof(f6613, plain, (~ (op2(e24, e25) = h6(op1(e14, e15))) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6612, f4797])).
fof(f6612, plain, (~ (h6(op1(e14, e15)) = op2(h6(e14), e25)) | ~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6611, f1243])).
fof(f6611, plain, (~ (h6(op1(e14, e16)) = op2(e24, h6(e16))) | ~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6610, f4797])).
fof(f6610, plain, (~ (op2(e25, e20) = h6(op1(e15, e10))) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6609, f1243])).
fof(f6609, plain, (~ (h6(op1(e15, e10)) = op2(h6(e15), e20)) | ~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6608, f4820])).
fof(f6608, plain, (~ (e24 = h6(op1(e15, e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6607, f4816])).
fof(f6607, plain, (~ (h6(op1(e15, e11)) = op2(e25, h6(e11))) | ~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6606, f1243])).
fof(f6606, plain, (~ (h6(op1(e15, e12)) = op2(e25, h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6605, f1243])).
fof(f6605, plain, (~ (h6(op1(e15, e12)) = op2(h6(e15), h5(e11))) | ~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6604, f4802])).
fof(f6604, plain, (~ (op2(e25, e23) = h6(op1(e15, e13))) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6603, f1243])).
fof(f6603, plain, (~ (h6(op1(e15, e13)) = op2(h6(e15), e23)) | ~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6602, f4800])).
fof(f6602, plain, (~ (op2(e25, e24) = h6(op1(e15, e14))) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6601, f1243])).
fof(f6601, plain, (~ (h6(op1(e15, e14)) = op2(h6(e15), e24)) | ~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6600, f4797])).
fof(f6600, plain, (~ (h6(e11) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6599, f1245])).
fof(f6599, plain, (~ (op2(e25, e25) = h6(op1(e15, e15))) | ~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6598, f1243])).
fof(f6598, plain, (~ (h6(op1(e15, e16)) = op2(e25, h6(e16))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6597, f1243])).
fof(f6597, plain, (~ (h6(op1(e16, e10)) = op2(h6(e16), e20)) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6596, f4820])).
fof(f6596, plain, (~ (h6(op1(e16, e12)) = op2(h6(e16), h5(e11))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) | ~ (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6595, f4802])).
fof(f6595, plain, (~ (h6(op1(e16, e13)) = op2(h6(e16), e23)) | ~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) | ~ (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6594, f4800])).
fof(f6594, plain, (~ (h6(op1(e16, e14)) = op2(h6(e16), e24)) | ~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) | ~ (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f6593, f4797])).
fof(f6593, plain, (~ (h6(op1(e16, e15)) = op2(h6(e16), e25)) | ~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e14)) = op2(h6(e16), h6(e14))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) | ~ (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(forward_demodulation, [], [f1592, f1243])).
fof(f1592, plain, (~ (e26 = h6(e16)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e15)) = op2(h6(e16), h6(e15))) | ~ (h6(op1(e16, e14)) = op2(h6(e16), h6(e14))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) | ~ (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))), inference(cnf_transformation, [], [f68])).
fof(f68, plain, (((~ (e26 = h7(e16)) & ~ (e26 = h7(e15)) & ~ (e26 = h7(e14)) & ~ (e26 = h7(e13)) & ~ (e26 = h7(e12)) & ~ (e26 = h7(e11)) & ~ (e26 = h7(e10))) | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | ~ (h7(op1(e16, e16)) = op2(h7(e16), h7(e16))) | ~ (h7(op1(e16, e15)) = op2(h7(e16), h7(e15))) | ~ (h7(op1(e16, e14)) = op2(h7(e16), h7(e14))) | ~ (h7(op1(e16, e13)) = op2(h7(e16), h7(e13))) | ~ (h7(op1(e16, e12)) = op2(h7(e16), h7(e12))) | ~ (h7(op1(e16, e11)) = op2(h7(e16), h7(e11))) | ~ (h7(op1(e16, e10)) = op2(h7(e16), h7(e10))) | ~ (h7(op1(e15, e16)) = op2(h7(e15), h7(e16))) | ~ (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) | ~ (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) | ~ (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) | ~ (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) | ~ (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) | ~ (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) | ~ (h7(op1(e14, e16)) = op2(h7(e14), h7(e16))) | ~ (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e16)) = op2(h7(e13), h7(e16))) | ~ (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e16)) = op2(h7(e12), h7(e16))) | ~ (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e16)) = op2(h7(e11), h7(e16))) | ~ (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e16)) = op2(h7(e10), h7(e16))) | ~ (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e26 = h6(e16)) & ~ (e26 = h6(e15)) & ~ (e26 = h6(e14)) & ~ (e26 = h6(e13)) & ~ (e26 = h6(e12)) & ~ (e26 = h6(e11)) & ~ (e26 = h6(e10))) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e15)) = op2(h6(e16), h6(e15))) | ~ (h6(op1(e16, e14)) = op2(h6(e16), h6(e14))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) | ~ (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e26 = h5(e16)) & ~ (e26 = h5(e15)) & ~ (e26 = h5(e14)) & ~ (e26 = h5(e13)) & ~ (e26 = h5(e12)) & ~ (e26 = h5(e11)) & ~ (e26 = h5(e10))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | ~ (h5(op1(e16, e16)) = op2(h5(e16), h5(e16))) | ~ (h5(op1(e16, e15)) = op2(h5(e16), h5(e15))) | ~ (h5(op1(e16, e14)) = op2(h5(e16), h5(e14))) | ~ (h5(op1(e16, e13)) = op2(h5(e16), h5(e13))) | ~ (h5(op1(e16, e12)) = op2(h5(e16), h5(e12))) | ~ (h5(op1(e16, e11)) = op2(h5(e16), h5(e11))) | ~ (h5(op1(e16, e10)) = op2(h5(e16), h5(e10))) | ~ (h5(op1(e15, e16)) = op2(h5(e15), h5(e16))) | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e16)) = op2(h5(e14), h5(e16))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e16)) = op2(h5(e13), h5(e16))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e16)) = op2(h5(e12), h5(e16))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e16)) = op2(h5(e11), h5(e16))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e16)) = op2(h5(e10), h5(e16))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e26 = h4(e16)) & ~ (e26 = h4(e15)) & ~ (e26 = h4(e14)) & ~ (e26 = h4(e13)) & ~ (e26 = h4(e12)) & ~ (e26 = h4(e11)) & ~ (e26 = h4(e10))) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | ~ (h4(op1(e16, e16)) = op2(h4(e16), h4(e16))) | ~ (h4(op1(e16, e15)) = op2(h4(e16), h4(e15))) | ~ (h4(op1(e16, e14)) = op2(h4(e16), h4(e14))) | ~ (h4(op1(e16, e13)) = op2(h4(e16), h4(e13))) | ~ (h4(op1(e16, e12)) = op2(h4(e16), h4(e12))) | ~ (h4(op1(e16, e11)) = op2(h4(e16), h4(e11))) | ~ (h4(op1(e16, e10)) = op2(h4(e16), h4(e10))) | ~ (h4(op1(e15, e16)) = op2(h4(e15), h4(e16))) | ~ (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) | ~ (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) | ~ (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) | ~ (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) | ~ (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) | ~ (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) | ~ (h4(op1(e14, e16)) = op2(h4(e14), h4(e16))) | ~ (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e16)) = op2(h4(e13), h4(e16))) | ~ (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e16)) = op2(h4(e12), h4(e16))) | ~ (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e16)) = op2(h4(e11), h4(e16))) | ~ (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e16)) = op2(h4(e10), h4(e16))) | ~ (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e26 = h3(e16)) & ~ (e26 = h3(e15)) & ~ (e26 = h3(e14)) & ~ (e26 = h3(e13)) & ~ (e26 = h3(e12)) & ~ (e26 = h3(e11)) & ~ (e26 = h3(e10))) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | ~ (h3(op1(e16, e16)) = op2(h3(e16), h3(e16))) | ~ (h3(op1(e16, e15)) = op2(h3(e16), h3(e15))) | ~ (h3(op1(e16, e14)) = op2(h3(e16), h3(e14))) | ~ (h3(op1(e16, e13)) = op2(h3(e16), h3(e13))) | ~ (h3(op1(e16, e12)) = op2(h3(e16), h3(e12))) | ~ (h3(op1(e16, e11)) = op2(h3(e16), h3(e11))) | ~ (h3(op1(e16, e10)) = op2(h3(e16), h3(e10))) | ~ (h3(op1(e15, e16)) = op2(h3(e15), h3(e16))) | ~ (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) | ~ (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) | ~ (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) | ~ (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) | ~ (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) | ~ (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) | ~ (h3(op1(e14, e16)) = op2(h3(e14), h3(e16))) | ~ (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e16)) = op2(h3(e13), h3(e16))) | ~ (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e16)) = op2(h3(e12), h3(e16))) | ~ (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e16)) = op2(h3(e11), h3(e16))) | ~ (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e16)) = op2(h3(e10), h3(e16))) | ~ (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e26 = h2(e16)) & ~ (e26 = h2(e15)) & ~ (e26 = h2(e14)) & ~ (e26 = h2(e13)) & ~ (e26 = h2(e12)) & ~ (e26 = h2(e11)) & ~ (e26 = h2(e10))) | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | ~ (h2(op1(e16, e16)) = op2(h2(e16), h2(e16))) | ~ (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) | ~ (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e26 = h1(e16)) & ~ (e26 = h1(e15)) & ~ (e26 = h1(e14)) & ~ (e26 = h1(e13)) & ~ (e26 = h1(e12)) & ~ (e26 = h1(e11)) & ~ (e26 = h1(e10))) | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (h1(op1(e16, e16)) = op2(h1(e16), h1(e16))) | ~ (h1(op1(e16, e15)) = op2(h1(e16), h1(e15))) | ~ (h1(op1(e16, e14)) = op2(h1(e16), h1(e14))) | ~ (h1(op1(e16, e13)) = op2(h1(e16), h1(e13))) | ~ (h1(op1(e16, e12)) = op2(h1(e16), h1(e12))) | ~ (h1(op1(e16, e11)) = op2(h1(e16), h1(e11))) | ~ (h1(op1(e16, e10)) = op2(h1(e16), h1(e10))) | ~ (h1(op1(e15, e16)) = op2(h1(e15), h1(e16))) | ~ (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) | ~ (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) | ~ (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) | ~ (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) | ~ (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) | ~ (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) | ~ (h1(op1(e14, e16)) = op2(h1(e14), h1(e16))) | ~ (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e16)) = op2(h1(e13), h1(e16))) | ~ (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e16)) = op2(h1(e12), h1(e16))) | ~ (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e16)) = op2(h1(e11), h1(e16))) | ~ (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e16)) = op2(h1(e10), h1(e16))) | ~ (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f25, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26])).
fof(f26, plain, ((~ (e20 = h1(e16)) & ~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(usedef, [], [e26])).
fof(e26, plain, (sP0 <=> (~ (e20 = h1(e16)) & ~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f27, plain, ((~ (e21 = h1(e16)) & ~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(usedef, [], [e27])).
fof(e27, plain, (sP1 <=> (~ (e21 = h1(e16)) & ~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f28, plain, ((~ (e22 = h1(e16)) & ~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(usedef, [], [e28])).
fof(e28, plain, (sP2 <=> (~ (e22 = h1(e16)) & ~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f29, plain, ((~ (e23 = h1(e16)) & ~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | ~ sP3), inference(usedef, [], [e29])).
fof(e29, plain, (sP3 <=> (~ (e23 = h1(e16)) & ~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f30, plain, ((~ (e24 = h1(e16)) & ~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | ~ sP4), inference(usedef, [], [e30])).
fof(e30, plain, (sP4 <=> (~ (e24 = h1(e16)) & ~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f31, plain, ((~ (e25 = h1(e16)) & ~ (e25 = h1(e15)) & ~ (e25 = h1(e14)) & ~ (e25 = h1(e13)) & ~ (e25 = h1(e12)) & ~ (e25 = h1(e11)) & ~ (e25 = h1(e10))) | ~ sP5), inference(usedef, [], [e31])).
fof(e31, plain, (sP5 <=> (~ (e25 = h1(e16)) & ~ (e25 = h1(e15)) & ~ (e25 = h1(e14)) & ~ (e25 = h1(e13)) & ~ (e25 = h1(e12)) & ~ (e25 = h1(e11)) & ~ (e25 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f32, plain, ((~ (e20 = h2(e16)) & ~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP6), inference(usedef, [], [e32])).
fof(e32, plain, (sP6 <=> (~ (e20 = h2(e16)) & ~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f33, plain, ((~ (e21 = h2(e16)) & ~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP7), inference(usedef, [], [e33])).
fof(e33, plain, (sP7 <=> (~ (e21 = h2(e16)) & ~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f34, plain, ((~ (e22 = h2(e16)) & ~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP8), inference(usedef, [], [e34])).
fof(e34, plain, (sP8 <=> (~ (e22 = h2(e16)) & ~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f35, plain, ((~ (e23 = h2(e16)) & ~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP9), inference(usedef, [], [e35])).
fof(e35, plain, (sP9 <=> (~ (e23 = h2(e16)) & ~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f36, plain, ((~ (e24 = h2(e16)) & ~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | ~ sP10), inference(usedef, [], [e36])).
fof(e36, plain, (sP10 <=> (~ (e24 = h2(e16)) & ~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f37, plain, ((~ (e25 = h2(e16)) & ~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10))) | ~ sP11), inference(usedef, [], [e37])).
fof(e37, plain, (sP11 <=> (~ (e25 = h2(e16)) & ~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f38, plain, ((~ (e20 = h3(e16)) & ~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP12), inference(usedef, [], [e38])).
fof(e38, plain, (sP12 <=> (~ (e20 = h3(e16)) & ~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f39, plain, ((~ (e21 = h3(e16)) & ~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP13), inference(usedef, [], [e39])).
fof(e39, plain, (sP13 <=> (~ (e21 = h3(e16)) & ~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f40, plain, ((~ (e22 = h3(e16)) & ~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP14), inference(usedef, [], [e40])).
fof(e40, plain, (sP14 <=> (~ (e22 = h3(e16)) & ~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f41, plain, ((~ (e23 = h3(e16)) & ~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP15), inference(usedef, [], [e41])).
fof(e41, plain, (sP15 <=> (~ (e23 = h3(e16)) & ~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f42, plain, ((~ (e24 = h3(e16)) & ~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | ~ sP16), inference(usedef, [], [e42])).
fof(e42, plain, (sP16 <=> (~ (e24 = h3(e16)) & ~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f43, plain, ((~ (e25 = h3(e16)) & ~ (e25 = h3(e15)) & ~ (e25 = h3(e14)) & ~ (e25 = h3(e13)) & ~ (e25 = h3(e12)) & ~ (e25 = h3(e11)) & ~ (e25 = h3(e10))) | ~ sP17), inference(usedef, [], [e43])).
fof(e43, plain, (sP17 <=> (~ (e25 = h3(e16)) & ~ (e25 = h3(e15)) & ~ (e25 = h3(e14)) & ~ (e25 = h3(e13)) & ~ (e25 = h3(e12)) & ~ (e25 = h3(e11)) & ~ (e25 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f44, plain, ((~ (e20 = h4(e16)) & ~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP18), inference(usedef, [], [e44])).
fof(e44, plain, (sP18 <=> (~ (e20 = h4(e16)) & ~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f45, plain, ((~ (e21 = h4(e16)) & ~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP19), inference(usedef, [], [e45])).
fof(e45, plain, (sP19 <=> (~ (e21 = h4(e16)) & ~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f46, plain, ((~ (e22 = h4(e16)) & ~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP20), inference(usedef, [], [e46])).
fof(e46, plain, (sP20 <=> (~ (e22 = h4(e16)) & ~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f47, plain, ((~ (e23 = h4(e16)) & ~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | ~ sP21), inference(usedef, [], [e47])).
fof(e47, plain, (sP21 <=> (~ (e23 = h4(e16)) & ~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f48, plain, ((~ (e24 = h4(e16)) & ~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | ~ sP22), inference(usedef, [], [e48])).
fof(e48, plain, (sP22 <=> (~ (e24 = h4(e16)) & ~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f49, plain, ((~ (e25 = h4(e16)) & ~ (e25 = h4(e15)) & ~ (e25 = h4(e14)) & ~ (e25 = h4(e13)) & ~ (e25 = h4(e12)) & ~ (e25 = h4(e11)) & ~ (e25 = h4(e10))) | ~ sP23), inference(usedef, [], [e49])).
fof(e49, plain, (sP23 <=> (~ (e25 = h4(e16)) & ~ (e25 = h4(e15)) & ~ (e25 = h4(e14)) & ~ (e25 = h4(e13)) & ~ (e25 = h4(e12)) & ~ (e25 = h4(e11)) & ~ (e25 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f50, plain, ((~ (e20 = h5(e16)) & ~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP24), inference(usedef, [], [e50])).
fof(e50, plain, (sP24 <=> (~ (e20 = h5(e16)) & ~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f51, plain, ((~ (e21 = h5(e16)) & ~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP25), inference(usedef, [], [e51])).
fof(e51, plain, (sP25 <=> (~ (e21 = h5(e16)) & ~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f52, plain, ((~ (e22 = h5(e16)) & ~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP26), inference(usedef, [], [e52])).
fof(e52, plain, (sP26 <=> (~ (e22 = h5(e16)) & ~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f53, plain, ((~ (e23 = h5(e16)) & ~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP27), inference(usedef, [], [e53])).
fof(e53, plain, (sP27 <=> (~ (e23 = h5(e16)) & ~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f54, plain, ((~ (e24 = h5(e16)) & ~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | ~ sP28), inference(usedef, [], [e54])).
fof(e54, plain, (sP28 <=> (~ (e24 = h5(e16)) & ~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f55, plain, ((~ (e25 = h5(e16)) & ~ (e25 = h5(e15)) & ~ (e25 = h5(e14)) & ~ (e25 = h5(e13)) & ~ (e25 = h5(e12)) & ~ (e25 = h5(e11)) & ~ (e25 = h5(e10))) | ~ sP29), inference(usedef, [], [e55])).
fof(e55, plain, (sP29 <=> (~ (e25 = h5(e16)) & ~ (e25 = h5(e15)) & ~ (e25 = h5(e14)) & ~ (e25 = h5(e13)) & ~ (e25 = h5(e12)) & ~ (e25 = h5(e11)) & ~ (e25 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f56, plain, ((~ (e20 = h6(e16)) & ~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP30), inference(usedef, [], [e56])).
fof(e56, plain, (sP30 <=> (~ (e20 = h6(e16)) & ~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f57, plain, ((~ (e21 = h6(e16)) & ~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP31), inference(usedef, [], [e57])).
fof(e57, plain, (sP31 <=> (~ (e21 = h6(e16)) & ~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f58, plain, ((~ (e22 = h6(e16)) & ~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP32), inference(usedef, [], [e58])).
fof(e58, plain, (sP32 <=> (~ (e22 = h6(e16)) & ~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f59, plain, ((~ (e23 = h6(e16)) & ~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | ~ sP33), inference(usedef, [], [e59])).
fof(e59, plain, (sP33 <=> (~ (e23 = h6(e16)) & ~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f60, plain, ((~ (e24 = h6(e16)) & ~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | ~ sP34), inference(usedef, [], [e60])).
fof(e60, plain, (sP34 <=> (~ (e24 = h6(e16)) & ~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f61, plain, ((~ (e25 = h6(e16)) & ~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10))) | ~ sP35), inference(usedef, [], [e61])).
fof(e61, plain, (sP35 <=> (~ (e25 = h6(e16)) & ~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f62, plain, ((~ (e20 = h7(e16)) & ~ (e20 = h7(e15)) & ~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP36), inference(usedef, [], [e62])).
fof(e62, plain, (sP36 <=> (~ (e20 = h7(e16)) & ~ (e20 = h7(e15)) & ~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f63, plain, ((~ (e21 = h7(e16)) & ~ (e21 = h7(e15)) & ~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP37), inference(usedef, [], [e63])).
fof(e63, plain, (sP37 <=> (~ (e21 = h7(e16)) & ~ (e21 = h7(e15)) & ~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f64, plain, ((~ (e22 = h7(e16)) & ~ (e22 = h7(e15)) & ~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP38), inference(usedef, [], [e64])).
fof(e64, plain, (sP38 <=> (~ (e22 = h7(e16)) & ~ (e22 = h7(e15)) & ~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f65, plain, ((~ (e23 = h7(e16)) & ~ (e23 = h7(e15)) & ~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | ~ sP39), inference(usedef, [], [e65])).
fof(e65, plain, (sP39 <=> (~ (e23 = h7(e16)) & ~ (e23 = h7(e15)) & ~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f66, plain, ((~ (e24 = h7(e16)) & ~ (e24 = h7(e15)) & ~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | ~ sP40), inference(usedef, [], [e66])).
fof(e66, plain, (sP40 <=> (~ (e24 = h7(e16)) & ~ (e24 = h7(e15)) & ~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f67, plain, ((~ (e25 = h7(e16)) & ~ (e25 = h7(e15)) & ~ (e25 = h7(e14)) & ~ (e25 = h7(e13)) & ~ (e25 = h7(e12)) & ~ (e25 = h7(e11)) & ~ (e25 = h7(e10))) | ~ sP41), inference(usedef, [], [e67])).
fof(e67, plain, (sP41 <=> (~ (e25 = h7(e16)) & ~ (e25 = h7(e15)) & ~ (e25 = h7(e14)) & ~ (e25 = h7(e13)) & ~ (e25 = h7(e12)) & ~ (e25 = h7(e11)) & ~ (e25 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f25, plain, (((~ (e26 = h7(e16)) & ~ (e26 = h7(e15)) & ~ (e26 = h7(e14)) & ~ (e26 = h7(e13)) & ~ (e26 = h7(e12)) & ~ (e26 = h7(e11)) & ~ (e26 = h7(e10))) | (~ (e25 = h7(e16)) & ~ (e25 = h7(e15)) & ~ (e25 = h7(e14)) & ~ (e25 = h7(e13)) & ~ (e25 = h7(e12)) & ~ (e25 = h7(e11)) & ~ (e25 = h7(e10))) | (~ (e24 = h7(e16)) & ~ (e24 = h7(e15)) & ~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | (~ (e23 = h7(e16)) & ~ (e23 = h7(e15)) & ~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | (~ (e22 = h7(e16)) & ~ (e22 = h7(e15)) & ~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | (~ (e21 = h7(e16)) & ~ (e21 = h7(e15)) & ~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | (~ (e20 = h7(e16)) & ~ (e20 = h7(e15)) & ~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ (h7(op1(e16, e16)) = op2(h7(e16), h7(e16))) | ~ (h7(op1(e16, e15)) = op2(h7(e16), h7(e15))) | ~ (h7(op1(e16, e14)) = op2(h7(e16), h7(e14))) | ~ (h7(op1(e16, e13)) = op2(h7(e16), h7(e13))) | ~ (h7(op1(e16, e12)) = op2(h7(e16), h7(e12))) | ~ (h7(op1(e16, e11)) = op2(h7(e16), h7(e11))) | ~ (h7(op1(e16, e10)) = op2(h7(e16), h7(e10))) | ~ (h7(op1(e15, e16)) = op2(h7(e15), h7(e16))) | ~ (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) | ~ (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) | ~ (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) | ~ (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) | ~ (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) | ~ (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) | ~ (h7(op1(e14, e16)) = op2(h7(e14), h7(e16))) | ~ (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e16)) = op2(h7(e13), h7(e16))) | ~ (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e16)) = op2(h7(e12), h7(e16))) | ~ (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e16)) = op2(h7(e11), h7(e16))) | ~ (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e16)) = op2(h7(e10), h7(e16))) | ~ (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e26 = h6(e16)) & ~ (e26 = h6(e15)) & ~ (e26 = h6(e14)) & ~ (e26 = h6(e13)) & ~ (e26 = h6(e12)) & ~ (e26 = h6(e11)) & ~ (e26 = h6(e10))) | (~ (e25 = h6(e16)) & ~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10))) | (~ (e24 = h6(e16)) & ~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | (~ (e23 = h6(e16)) & ~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e16)) & ~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e16)) & ~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e16)) & ~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) | ~ (h6(op1(e16, e15)) = op2(h6(e16), h6(e15))) | ~ (h6(op1(e16, e14)) = op2(h6(e16), h6(e14))) | ~ (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) | ~ (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) | ~ (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) | ~ (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) | ~ (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) | ~ (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) | ~ (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) | ~ (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) | ~ (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) | ~ (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) | ~ (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) | ~ (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) | ~ (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) | ~ (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) | ~ (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) | ~ (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) | ~ (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e26 = h5(e16)) & ~ (e26 = h5(e15)) & ~ (e26 = h5(e14)) & ~ (e26 = h5(e13)) & ~ (e26 = h5(e12)) & ~ (e26 = h5(e11)) & ~ (e26 = h5(e10))) | (~ (e25 = h5(e16)) & ~ (e25 = h5(e15)) & ~ (e25 = h5(e14)) & ~ (e25 = h5(e13)) & ~ (e25 = h5(e12)) & ~ (e25 = h5(e11)) & ~ (e25 = h5(e10))) | (~ (e24 = h5(e16)) & ~ (e24 = h5(e15)) & ~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | (~ (e23 = h5(e16)) & ~ (e23 = h5(e15)) & ~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e16)) & ~ (e22 = h5(e15)) & ~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e16)) & ~ (e21 = h5(e15)) & ~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e16)) & ~ (e20 = h5(e15)) & ~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e16, e16)) = op2(h5(e16), h5(e16))) | ~ (h5(op1(e16, e15)) = op2(h5(e16), h5(e15))) | ~ (h5(op1(e16, e14)) = op2(h5(e16), h5(e14))) | ~ (h5(op1(e16, e13)) = op2(h5(e16), h5(e13))) | ~ (h5(op1(e16, e12)) = op2(h5(e16), h5(e12))) | ~ (h5(op1(e16, e11)) = op2(h5(e16), h5(e11))) | ~ (h5(op1(e16, e10)) = op2(h5(e16), h5(e10))) | ~ (h5(op1(e15, e16)) = op2(h5(e15), h5(e16))) | ~ (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) | ~ (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) | ~ (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) | ~ (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) | ~ (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) | ~ (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) | ~ (h5(op1(e14, e16)) = op2(h5(e14), h5(e16))) | ~ (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e16)) = op2(h5(e13), h5(e16))) | ~ (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e16)) = op2(h5(e12), h5(e16))) | ~ (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e16)) = op2(h5(e11), h5(e16))) | ~ (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e16)) = op2(h5(e10), h5(e16))) | ~ (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e26 = h4(e16)) & ~ (e26 = h4(e15)) & ~ (e26 = h4(e14)) & ~ (e26 = h4(e13)) & ~ (e26 = h4(e12)) & ~ (e26 = h4(e11)) & ~ (e26 = h4(e10))) | (~ (e25 = h4(e16)) & ~ (e25 = h4(e15)) & ~ (e25 = h4(e14)) & ~ (e25 = h4(e13)) & ~ (e25 = h4(e12)) & ~ (e25 = h4(e11)) & ~ (e25 = h4(e10))) | (~ (e24 = h4(e16)) & ~ (e24 = h4(e15)) & ~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | (~ (e23 = h4(e16)) & ~ (e23 = h4(e15)) & ~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e16)) & ~ (e22 = h4(e15)) & ~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e16)) & ~ (e21 = h4(e15)) & ~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e16)) & ~ (e20 = h4(e15)) & ~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e16, e16)) = op2(h4(e16), h4(e16))) | ~ (h4(op1(e16, e15)) = op2(h4(e16), h4(e15))) | ~ (h4(op1(e16, e14)) = op2(h4(e16), h4(e14))) | ~ (h4(op1(e16, e13)) = op2(h4(e16), h4(e13))) | ~ (h4(op1(e16, e12)) = op2(h4(e16), h4(e12))) | ~ (h4(op1(e16, e11)) = op2(h4(e16), h4(e11))) | ~ (h4(op1(e16, e10)) = op2(h4(e16), h4(e10))) | ~ (h4(op1(e15, e16)) = op2(h4(e15), h4(e16))) | ~ (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) | ~ (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) | ~ (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) | ~ (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) | ~ (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) | ~ (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) | ~ (h4(op1(e14, e16)) = op2(h4(e14), h4(e16))) | ~ (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e16)) = op2(h4(e13), h4(e16))) | ~ (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e16)) = op2(h4(e12), h4(e16))) | ~ (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e16)) = op2(h4(e11), h4(e16))) | ~ (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e16)) = op2(h4(e10), h4(e16))) | ~ (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e26 = h3(e16)) & ~ (e26 = h3(e15)) & ~ (e26 = h3(e14)) & ~ (e26 = h3(e13)) & ~ (e26 = h3(e12)) & ~ (e26 = h3(e11)) & ~ (e26 = h3(e10))) | (~ (e25 = h3(e16)) & ~ (e25 = h3(e15)) & ~ (e25 = h3(e14)) & ~ (e25 = h3(e13)) & ~ (e25 = h3(e12)) & ~ (e25 = h3(e11)) & ~ (e25 = h3(e10))) | (~ (e24 = h3(e16)) & ~ (e24 = h3(e15)) & ~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | (~ (e23 = h3(e16)) & ~ (e23 = h3(e15)) & ~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e16)) & ~ (e22 = h3(e15)) & ~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e16)) & ~ (e21 = h3(e15)) & ~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e16)) & ~ (e20 = h3(e15)) & ~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e16, e16)) = op2(h3(e16), h3(e16))) | ~ (h3(op1(e16, e15)) = op2(h3(e16), h3(e15))) | ~ (h3(op1(e16, e14)) = op2(h3(e16), h3(e14))) | ~ (h3(op1(e16, e13)) = op2(h3(e16), h3(e13))) | ~ (h3(op1(e16, e12)) = op2(h3(e16), h3(e12))) | ~ (h3(op1(e16, e11)) = op2(h3(e16), h3(e11))) | ~ (h3(op1(e16, e10)) = op2(h3(e16), h3(e10))) | ~ (h3(op1(e15, e16)) = op2(h3(e15), h3(e16))) | ~ (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) | ~ (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) | ~ (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) | ~ (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) | ~ (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) | ~ (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) | ~ (h3(op1(e14, e16)) = op2(h3(e14), h3(e16))) | ~ (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e16)) = op2(h3(e13), h3(e16))) | ~ (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e16)) = op2(h3(e12), h3(e16))) | ~ (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e16)) = op2(h3(e11), h3(e16))) | ~ (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e16)) = op2(h3(e10), h3(e16))) | ~ (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e26 = h2(e16)) & ~ (e26 = h2(e15)) & ~ (e26 = h2(e14)) & ~ (e26 = h2(e13)) & ~ (e26 = h2(e12)) & ~ (e26 = h2(e11)) & ~ (e26 = h2(e10))) | (~ (e25 = h2(e16)) & ~ (e25 = h2(e15)) & ~ (e25 = h2(e14)) & ~ (e25 = h2(e13)) & ~ (e25 = h2(e12)) & ~ (e25 = h2(e11)) & ~ (e25 = h2(e10))) | (~ (e24 = h2(e16)) & ~ (e24 = h2(e15)) & ~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | (~ (e23 = h2(e16)) & ~ (e23 = h2(e15)) & ~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e16)) & ~ (e22 = h2(e15)) & ~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e16)) & ~ (e21 = h2(e15)) & ~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e16)) & ~ (e20 = h2(e15)) & ~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e16, e16)) = op2(h2(e16), h2(e16))) | ~ (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) | ~ (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) | ~ (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) | ~ (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) | ~ (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) | ~ (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) | ~ (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) | ~ (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) | ~ (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) | ~ (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) | ~ (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) | ~ (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) | ~ (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) | ~ (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) | ~ (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) | ~ (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) | ~ (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) | ~ (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) | ~ (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e26 = h1(e16)) & ~ (e26 = h1(e15)) & ~ (e26 = h1(e14)) & ~ (e26 = h1(e13)) & ~ (e26 = h1(e12)) & ~ (e26 = h1(e11)) & ~ (e26 = h1(e10))) | (~ (e25 = h1(e16)) & ~ (e25 = h1(e15)) & ~ (e25 = h1(e14)) & ~ (e25 = h1(e13)) & ~ (e25 = h1(e12)) & ~ (e25 = h1(e11)) & ~ (e25 = h1(e10))) | (~ (e24 = h1(e16)) & ~ (e24 = h1(e15)) & ~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | (~ (e23 = h1(e16)) & ~ (e23 = h1(e15)) & ~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e16)) & ~ (e22 = h1(e15)) & ~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e16)) & ~ (e21 = h1(e15)) & ~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e16)) & ~ (e20 = h1(e15)) & ~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e16, e16)) = op2(h1(e16), h1(e16))) | ~ (h1(op1(e16, e15)) = op2(h1(e16), h1(e15))) | ~ (h1(op1(e16, e14)) = op2(h1(e16), h1(e14))) | ~ (h1(op1(e16, e13)) = op2(h1(e16), h1(e13))) | ~ (h1(op1(e16, e12)) = op2(h1(e16), h1(e12))) | ~ (h1(op1(e16, e11)) = op2(h1(e16), h1(e11))) | ~ (h1(op1(e16, e10)) = op2(h1(e16), h1(e10))) | ~ (h1(op1(e15, e16)) = op2(h1(e15), h1(e16))) | ~ (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) | ~ (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) | ~ (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) | ~ (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) | ~ (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) | ~ (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) | ~ (h1(op1(e14, e16)) = op2(h1(e14), h1(e16))) | ~ (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e16)) = op2(h1(e13), h1(e16))) | ~ (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e16)) = op2(h1(e12), h1(e16))) | ~ (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e16)) = op2(h1(e11), h1(e16))) | ~ (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e16)) = op2(h1(e10), h1(e16))) | ~ (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((((e26 = h7(e16)) | (e26 = h7(e15)) | (e26 = h7(e14)) | (e26 = h7(e13)) | (e26 = h7(e12)) | (e26 = h7(e11)) | (e26 = h7(e10))) & ((e25 = h7(e16)) | (e25 = h7(e15)) | (e25 = h7(e14)) | (e25 = h7(e13)) | (e25 = h7(e12)) | (e25 = h7(e11)) | (e25 = h7(e10))) & ((e24 = h7(e16)) | (e24 = h7(e15)) | (e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e16)) | (e23 = h7(e15)) | (e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e16)) | (e22 = h7(e15)) | (e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e16)) | (e21 = h7(e15)) | (e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e16)) | (e20 = h7(e15)) | (e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e16, e16)) = op2(h7(e16), h7(e16))) & (h7(op1(e16, e15)) = op2(h7(e16), h7(e15))) & (h7(op1(e16, e14)) = op2(h7(e16), h7(e14))) & (h7(op1(e16, e13)) = op2(h7(e16), h7(e13))) & (h7(op1(e16, e12)) = op2(h7(e16), h7(e12))) & (h7(op1(e16, e11)) = op2(h7(e16), h7(e11))) & (h7(op1(e16, e10)) = op2(h7(e16), h7(e10))) & (h7(op1(e15, e16)) = op2(h7(e15), h7(e16))) & (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) & (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) & (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) & (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) & (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) & (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) & (h7(op1(e14, e16)) = op2(h7(e14), h7(e16))) & (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e16)) = op2(h7(e13), h7(e16))) & (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e16)) = op2(h7(e12), h7(e16))) & (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e16)) = op2(h7(e11), h7(e16))) & (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e16)) = op2(h7(e10), h7(e16))) & (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e26 = h6(e16)) | (e26 = h6(e15)) | (e26 = h6(e14)) | (e26 = h6(e13)) | (e26 = h6(e12)) | (e26 = h6(e11)) | (e26 = h6(e10))) & ((e25 = h6(e16)) | (e25 = h6(e15)) | (e25 = h6(e14)) | (e25 = h6(e13)) | (e25 = h6(e12)) | (e25 = h6(e11)) | (e25 = h6(e10))) & ((e24 = h6(e16)) | (e24 = h6(e15)) | (e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e16)) | (e23 = h6(e15)) | (e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e16)) | (e22 = h6(e15)) | (e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e16)) | (e21 = h6(e15)) | (e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e16)) | (e20 = h6(e15)) | (e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) & (h6(op1(e16, e15)) = op2(h6(e16), h6(e15))) & (h6(op1(e16, e14)) = op2(h6(e16), h6(e14))) & (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) & (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) & (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) & (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) & (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) & (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) & (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) & (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) & (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) & (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) & (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) & (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) & (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) & (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) & (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) & (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) & (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e26 = h5(e16)) | (e26 = h5(e15)) | (e26 = h5(e14)) | (e26 = h5(e13)) | (e26 = h5(e12)) | (e26 = h5(e11)) | (e26 = h5(e10))) & ((e25 = h5(e16)) | (e25 = h5(e15)) | (e25 = h5(e14)) | (e25 = h5(e13)) | (e25 = h5(e12)) | (e25 = h5(e11)) | (e25 = h5(e10))) & ((e24 = h5(e16)) | (e24 = h5(e15)) | (e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e16)) | (e23 = h5(e15)) | (e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e16)) | (e22 = h5(e15)) | (e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e16)) | (e21 = h5(e15)) | (e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e16)) | (e20 = h5(e15)) | (e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e16, e16)) = op2(h5(e16), h5(e16))) & (h5(op1(e16, e15)) = op2(h5(e16), h5(e15))) & (h5(op1(e16, e14)) = op2(h5(e16), h5(e14))) & (h5(op1(e16, e13)) = op2(h5(e16), h5(e13))) & (h5(op1(e16, e12)) = op2(h5(e16), h5(e12))) & (h5(op1(e16, e11)) = op2(h5(e16), h5(e11))) & (h5(op1(e16, e10)) = op2(h5(e16), h5(e10))) & (h5(op1(e15, e16)) = op2(h5(e15), h5(e16))) & (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) & (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) & (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) & (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) & (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) & (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) & (h5(op1(e14, e16)) = op2(h5(e14), h5(e16))) & (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e16)) = op2(h5(e13), h5(e16))) & (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e16)) = op2(h5(e12), h5(e16))) & (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e16)) = op2(h5(e11), h5(e16))) & (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e16)) = op2(h5(e10), h5(e16))) & (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e26 = h4(e16)) | (e26 = h4(e15)) | (e26 = h4(e14)) | (e26 = h4(e13)) | (e26 = h4(e12)) | (e26 = h4(e11)) | (e26 = h4(e10))) & ((e25 = h4(e16)) | (e25 = h4(e15)) | (e25 = h4(e14)) | (e25 = h4(e13)) | (e25 = h4(e12)) | (e25 = h4(e11)) | (e25 = h4(e10))) & ((e24 = h4(e16)) | (e24 = h4(e15)) | (e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e16)) | (e23 = h4(e15)) | (e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e16)) | (e22 = h4(e15)) | (e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e16)) | (e21 = h4(e15)) | (e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e16)) | (e20 = h4(e15)) | (e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e16, e16)) = op2(h4(e16), h4(e16))) & (h4(op1(e16, e15)) = op2(h4(e16), h4(e15))) & (h4(op1(e16, e14)) = op2(h4(e16), h4(e14))) & (h4(op1(e16, e13)) = op2(h4(e16), h4(e13))) & (h4(op1(e16, e12)) = op2(h4(e16), h4(e12))) & (h4(op1(e16, e11)) = op2(h4(e16), h4(e11))) & (h4(op1(e16, e10)) = op2(h4(e16), h4(e10))) & (h4(op1(e15, e16)) = op2(h4(e15), h4(e16))) & (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) & (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) & (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) & (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) & (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) & (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) & (h4(op1(e14, e16)) = op2(h4(e14), h4(e16))) & (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e16)) = op2(h4(e13), h4(e16))) & (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e16)) = op2(h4(e12), h4(e16))) & (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e16)) = op2(h4(e11), h4(e16))) & (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e16)) = op2(h4(e10), h4(e16))) & (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e26 = h3(e16)) | (e26 = h3(e15)) | (e26 = h3(e14)) | (e26 = h3(e13)) | (e26 = h3(e12)) | (e26 = h3(e11)) | (e26 = h3(e10))) & ((e25 = h3(e16)) | (e25 = h3(e15)) | (e25 = h3(e14)) | (e25 = h3(e13)) | (e25 = h3(e12)) | (e25 = h3(e11)) | (e25 = h3(e10))) & ((e24 = h3(e16)) | (e24 = h3(e15)) | (e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e16)) | (e23 = h3(e15)) | (e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e16)) | (e22 = h3(e15)) | (e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e16)) | (e21 = h3(e15)) | (e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e16)) | (e20 = h3(e15)) | (e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e16, e16)) = op2(h3(e16), h3(e16))) & (h3(op1(e16, e15)) = op2(h3(e16), h3(e15))) & (h3(op1(e16, e14)) = op2(h3(e16), h3(e14))) & (h3(op1(e16, e13)) = op2(h3(e16), h3(e13))) & (h3(op1(e16, e12)) = op2(h3(e16), h3(e12))) & (h3(op1(e16, e11)) = op2(h3(e16), h3(e11))) & (h3(op1(e16, e10)) = op2(h3(e16), h3(e10))) & (h3(op1(e15, e16)) = op2(h3(e15), h3(e16))) & (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) & (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) & (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) & (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) & (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) & (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) & (h3(op1(e14, e16)) = op2(h3(e14), h3(e16))) & (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e16)) = op2(h3(e13), h3(e16))) & (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e16)) = op2(h3(e12), h3(e16))) & (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e16)) = op2(h3(e11), h3(e16))) & (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e16)) = op2(h3(e10), h3(e16))) & (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e26 = h2(e16)) | (e26 = h2(e15)) | (e26 = h2(e14)) | (e26 = h2(e13)) | (e26 = h2(e12)) | (e26 = h2(e11)) | (e26 = h2(e10))) & ((e25 = h2(e16)) | (e25 = h2(e15)) | (e25 = h2(e14)) | (e25 = h2(e13)) | (e25 = h2(e12)) | (e25 = h2(e11)) | (e25 = h2(e10))) & ((e24 = h2(e16)) | (e24 = h2(e15)) | (e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e16)) | (e23 = h2(e15)) | (e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e16)) | (e22 = h2(e15)) | (e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e16)) | (e21 = h2(e15)) | (e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e16)) | (e20 = h2(e15)) | (e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e16, e16)) = op2(h2(e16), h2(e16))) & (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) & (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) & (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) & (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) & (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) & (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) & (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) & (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) & (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) & (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) & (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) & (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) & (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) & (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) & (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) & (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) & (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) & (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) & (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e26 = h1(e16)) | (e26 = h1(e15)) | (e26 = h1(e14)) | (e26 = h1(e13)) | (e26 = h1(e12)) | (e26 = h1(e11)) | (e26 = h1(e10))) & ((e25 = h1(e16)) | (e25 = h1(e15)) | (e25 = h1(e14)) | (e25 = h1(e13)) | (e25 = h1(e12)) | (e25 = h1(e11)) | (e25 = h1(e10))) & ((e24 = h1(e16)) | (e24 = h1(e15)) | (e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e16)) | (e23 = h1(e15)) | (e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e16)) | (e22 = h1(e15)) | (e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e16)) | (e21 = h1(e15)) | (e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e16)) | (e20 = h1(e15)) | (e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e16, e16)) = op2(h1(e16), h1(e16))) & (h1(op1(e16, e15)) = op2(h1(e16), h1(e15))) & (h1(op1(e16, e14)) = op2(h1(e16), h1(e14))) & (h1(op1(e16, e13)) = op2(h1(e16), h1(e13))) & (h1(op1(e16, e12)) = op2(h1(e16), h1(e12))) & (h1(op1(e16, e11)) = op2(h1(e16), h1(e11))) & (h1(op1(e16, e10)) = op2(h1(e16), h1(e10))) & (h1(op1(e15, e16)) = op2(h1(e15), h1(e16))) & (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) & (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) & (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) & (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) & (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) & (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) & (h1(op1(e14, e16)) = op2(h1(e14), h1(e16))) & (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e16)) = op2(h1(e13), h1(e16))) & (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e16)) = op2(h1(e12), h1(e16))) & (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e16)) = op2(h1(e11), h1(e16))) & (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e16)) = op2(h1(e10), h1(e16))) & (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f23])).
fof(f23, plain, ~ ((((e26 = h7(e16)) | (e26 = h7(e15)) | (e26 = h7(e14)) | (e26 = h7(e13)) | (e26 = h7(e12)) | (e26 = h7(e11)) | (e26 = h7(e10))) & ((e25 = h7(e16)) | (e25 = h7(e15)) | (e25 = h7(e14)) | (e25 = h7(e13)) | (e25 = h7(e12)) | (e25 = h7(e11)) | (e25 = h7(e10))) & ((e24 = h7(e16)) | (e24 = h7(e15)) | (e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e16)) | (e23 = h7(e15)) | (e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e16)) | (e22 = h7(e15)) | (e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e16)) | (e21 = h7(e15)) | (e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e16)) | (e20 = h7(e15)) | (e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e16, e16)) = op2(h7(e16), h7(e16))) & (h7(op1(e16, e15)) = op2(h7(e16), h7(e15))) & (h7(op1(e16, e14)) = op2(h7(e16), h7(e14))) & (h7(op1(e16, e13)) = op2(h7(e16), h7(e13))) & (h7(op1(e16, e12)) = op2(h7(e16), h7(e12))) & (h7(op1(e16, e11)) = op2(h7(e16), h7(e11))) & (h7(op1(e16, e10)) = op2(h7(e16), h7(e10))) & (h7(op1(e15, e16)) = op2(h7(e15), h7(e16))) & (h7(op1(e15, e15)) = op2(h7(e15), h7(e15))) & (h7(op1(e15, e14)) = op2(h7(e15), h7(e14))) & (h7(op1(e15, e13)) = op2(h7(e15), h7(e13))) & (h7(op1(e15, e12)) = op2(h7(e15), h7(e12))) & (h7(op1(e15, e11)) = op2(h7(e15), h7(e11))) & (h7(op1(e15, e10)) = op2(h7(e15), h7(e10))) & (h7(op1(e14, e16)) = op2(h7(e14), h7(e16))) & (h7(op1(e14, e15)) = op2(h7(e14), h7(e15))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e16)) = op2(h7(e13), h7(e16))) & (h7(op1(e13, e15)) = op2(h7(e13), h7(e15))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e16)) = op2(h7(e12), h7(e16))) & (h7(op1(e12, e15)) = op2(h7(e12), h7(e15))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e16)) = op2(h7(e11), h7(e16))) & (h7(op1(e11, e15)) = op2(h7(e11), h7(e15))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e16)) = op2(h7(e10), h7(e16))) & (h7(op1(e10, e15)) = op2(h7(e10), h7(e15))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e26 = h6(e16)) | (e26 = h6(e15)) | (e26 = h6(e14)) | (e26 = h6(e13)) | (e26 = h6(e12)) | (e26 = h6(e11)) | (e26 = h6(e10))) & ((e25 = h6(e16)) | (e25 = h6(e15)) | (e25 = h6(e14)) | (e25 = h6(e13)) | (e25 = h6(e12)) | (e25 = h6(e11)) | (e25 = h6(e10))) & ((e24 = h6(e16)) | (e24 = h6(e15)) | (e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e16)) | (e23 = h6(e15)) | (e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e16)) | (e22 = h6(e15)) | (e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e16)) | (e21 = h6(e15)) | (e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e16)) | (e20 = h6(e15)) | (e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e16, e16)) = op2(h6(e16), h6(e16))) & (h6(op1(e16, e15)) = op2(h6(e16), h6(e15))) & (h6(op1(e16, e14)) = op2(h6(e16), h6(e14))) & (h6(op1(e16, e13)) = op2(h6(e16), h6(e13))) & (h6(op1(e16, e12)) = op2(h6(e16), h6(e12))) & (h6(op1(e16, e11)) = op2(h6(e16), h6(e11))) & (h6(op1(e16, e10)) = op2(h6(e16), h6(e10))) & (h6(op1(e15, e16)) = op2(h6(e15), h6(e16))) & (h6(op1(e15, e15)) = op2(h6(e15), h6(e15))) & (h6(op1(e15, e14)) = op2(h6(e15), h6(e14))) & (h6(op1(e15, e13)) = op2(h6(e15), h6(e13))) & (h6(op1(e15, e12)) = op2(h6(e15), h6(e12))) & (h6(op1(e15, e11)) = op2(h6(e15), h6(e11))) & (h6(op1(e15, e10)) = op2(h6(e15), h6(e10))) & (h6(op1(e14, e16)) = op2(h6(e14), h6(e16))) & (h6(op1(e14, e15)) = op2(h6(e14), h6(e15))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e16)) = op2(h6(e13), h6(e16))) & (h6(op1(e13, e15)) = op2(h6(e13), h6(e15))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e16)) = op2(h6(e12), h6(e16))) & (h6(op1(e12, e15)) = op2(h6(e12), h6(e15))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e16)) = op2(h6(e11), h6(e16))) & (h6(op1(e11, e15)) = op2(h6(e11), h6(e15))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e16)) = op2(h6(e10), h6(e16))) & (h6(op1(e10, e15)) = op2(h6(e10), h6(e15))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e26 = h5(e16)) | (e26 = h5(e15)) | (e26 = h5(e14)) | (e26 = h5(e13)) | (e26 = h5(e12)) | (e26 = h5(e11)) | (e26 = h5(e10))) & ((e25 = h5(e16)) | (e25 = h5(e15)) | (e25 = h5(e14)) | (e25 = h5(e13)) | (e25 = h5(e12)) | (e25 = h5(e11)) | (e25 = h5(e10))) & ((e24 = h5(e16)) | (e24 = h5(e15)) | (e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e16)) | (e23 = h5(e15)) | (e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e16)) | (e22 = h5(e15)) | (e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e16)) | (e21 = h5(e15)) | (e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e16)) | (e20 = h5(e15)) | (e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e16, e16)) = op2(h5(e16), h5(e16))) & (h5(op1(e16, e15)) = op2(h5(e16), h5(e15))) & (h5(op1(e16, e14)) = op2(h5(e16), h5(e14))) & (h5(op1(e16, e13)) = op2(h5(e16), h5(e13))) & (h5(op1(e16, e12)) = op2(h5(e16), h5(e12))) & (h5(op1(e16, e11)) = op2(h5(e16), h5(e11))) & (h5(op1(e16, e10)) = op2(h5(e16), h5(e10))) & (h5(op1(e15, e16)) = op2(h5(e15), h5(e16))) & (h5(op1(e15, e15)) = op2(h5(e15), h5(e15))) & (h5(op1(e15, e14)) = op2(h5(e15), h5(e14))) & (h5(op1(e15, e13)) = op2(h5(e15), h5(e13))) & (h5(op1(e15, e12)) = op2(h5(e15), h5(e12))) & (h5(op1(e15, e11)) = op2(h5(e15), h5(e11))) & (h5(op1(e15, e10)) = op2(h5(e15), h5(e10))) & (h5(op1(e14, e16)) = op2(h5(e14), h5(e16))) & (h5(op1(e14, e15)) = op2(h5(e14), h5(e15))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e16)) = op2(h5(e13), h5(e16))) & (h5(op1(e13, e15)) = op2(h5(e13), h5(e15))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e16)) = op2(h5(e12), h5(e16))) & (h5(op1(e12, e15)) = op2(h5(e12), h5(e15))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e16)) = op2(h5(e11), h5(e16))) & (h5(op1(e11, e15)) = op2(h5(e11), h5(e15))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e16)) = op2(h5(e10), h5(e16))) & (h5(op1(e10, e15)) = op2(h5(e10), h5(e15))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e26 = h4(e16)) | (e26 = h4(e15)) | (e26 = h4(e14)) | (e26 = h4(e13)) | (e26 = h4(e12)) | (e26 = h4(e11)) | (e26 = h4(e10))) & ((e25 = h4(e16)) | (e25 = h4(e15)) | (e25 = h4(e14)) | (e25 = h4(e13)) | (e25 = h4(e12)) | (e25 = h4(e11)) | (e25 = h4(e10))) & ((e24 = h4(e16)) | (e24 = h4(e15)) | (e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e16)) | (e23 = h4(e15)) | (e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e16)) | (e22 = h4(e15)) | (e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e16)) | (e21 = h4(e15)) | (e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e16)) | (e20 = h4(e15)) | (e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e16, e16)) = op2(h4(e16), h4(e16))) & (h4(op1(e16, e15)) = op2(h4(e16), h4(e15))) & (h4(op1(e16, e14)) = op2(h4(e16), h4(e14))) & (h4(op1(e16, e13)) = op2(h4(e16), h4(e13))) & (h4(op1(e16, e12)) = op2(h4(e16), h4(e12))) & (h4(op1(e16, e11)) = op2(h4(e16), h4(e11))) & (h4(op1(e16, e10)) = op2(h4(e16), h4(e10))) & (h4(op1(e15, e16)) = op2(h4(e15), h4(e16))) & (h4(op1(e15, e15)) = op2(h4(e15), h4(e15))) & (h4(op1(e15, e14)) = op2(h4(e15), h4(e14))) & (h4(op1(e15, e13)) = op2(h4(e15), h4(e13))) & (h4(op1(e15, e12)) = op2(h4(e15), h4(e12))) & (h4(op1(e15, e11)) = op2(h4(e15), h4(e11))) & (h4(op1(e15, e10)) = op2(h4(e15), h4(e10))) & (h4(op1(e14, e16)) = op2(h4(e14), h4(e16))) & (h4(op1(e14, e15)) = op2(h4(e14), h4(e15))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e16)) = op2(h4(e13), h4(e16))) & (h4(op1(e13, e15)) = op2(h4(e13), h4(e15))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e16)) = op2(h4(e12), h4(e16))) & (h4(op1(e12, e15)) = op2(h4(e12), h4(e15))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e16)) = op2(h4(e11), h4(e16))) & (h4(op1(e11, e15)) = op2(h4(e11), h4(e15))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e16)) = op2(h4(e10), h4(e16))) & (h4(op1(e10, e15)) = op2(h4(e10), h4(e15))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e26 = h3(e16)) | (e26 = h3(e15)) | (e26 = h3(e14)) | (e26 = h3(e13)) | (e26 = h3(e12)) | (e26 = h3(e11)) | (e26 = h3(e10))) & ((e25 = h3(e16)) | (e25 = h3(e15)) | (e25 = h3(e14)) | (e25 = h3(e13)) | (e25 = h3(e12)) | (e25 = h3(e11)) | (e25 = h3(e10))) & ((e24 = h3(e16)) | (e24 = h3(e15)) | (e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e16)) | (e23 = h3(e15)) | (e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e16)) | (e22 = h3(e15)) | (e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e16)) | (e21 = h3(e15)) | (e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e16)) | (e20 = h3(e15)) | (e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e16, e16)) = op2(h3(e16), h3(e16))) & (h3(op1(e16, e15)) = op2(h3(e16), h3(e15))) & (h3(op1(e16, e14)) = op2(h3(e16), h3(e14))) & (h3(op1(e16, e13)) = op2(h3(e16), h3(e13))) & (h3(op1(e16, e12)) = op2(h3(e16), h3(e12))) & (h3(op1(e16, e11)) = op2(h3(e16), h3(e11))) & (h3(op1(e16, e10)) = op2(h3(e16), h3(e10))) & (h3(op1(e15, e16)) = op2(h3(e15), h3(e16))) & (h3(op1(e15, e15)) = op2(h3(e15), h3(e15))) & (h3(op1(e15, e14)) = op2(h3(e15), h3(e14))) & (h3(op1(e15, e13)) = op2(h3(e15), h3(e13))) & (h3(op1(e15, e12)) = op2(h3(e15), h3(e12))) & (h3(op1(e15, e11)) = op2(h3(e15), h3(e11))) & (h3(op1(e15, e10)) = op2(h3(e15), h3(e10))) & (h3(op1(e14, e16)) = op2(h3(e14), h3(e16))) & (h3(op1(e14, e15)) = op2(h3(e14), h3(e15))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e16)) = op2(h3(e13), h3(e16))) & (h3(op1(e13, e15)) = op2(h3(e13), h3(e15))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e16)) = op2(h3(e12), h3(e16))) & (h3(op1(e12, e15)) = op2(h3(e12), h3(e15))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e16)) = op2(h3(e11), h3(e16))) & (h3(op1(e11, e15)) = op2(h3(e11), h3(e15))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e16)) = op2(h3(e10), h3(e16))) & (h3(op1(e10, e15)) = op2(h3(e10), h3(e15))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e26 = h2(e16)) | (e26 = h2(e15)) | (e26 = h2(e14)) | (e26 = h2(e13)) | (e26 = h2(e12)) | (e26 = h2(e11)) | (e26 = h2(e10))) & ((e25 = h2(e16)) | (e25 = h2(e15)) | (e25 = h2(e14)) | (e25 = h2(e13)) | (e25 = h2(e12)) | (e25 = h2(e11)) | (e25 = h2(e10))) & ((e24 = h2(e16)) | (e24 = h2(e15)) | (e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e16)) | (e23 = h2(e15)) | (e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e16)) | (e22 = h2(e15)) | (e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e16)) | (e21 = h2(e15)) | (e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e16)) | (e20 = h2(e15)) | (e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e16, e16)) = op2(h2(e16), h2(e16))) & (h2(op1(e16, e15)) = op2(h2(e16), h2(e15))) & (h2(op1(e16, e14)) = op2(h2(e16), h2(e14))) & (h2(op1(e16, e13)) = op2(h2(e16), h2(e13))) & (h2(op1(e16, e12)) = op2(h2(e16), h2(e12))) & (h2(op1(e16, e11)) = op2(h2(e16), h2(e11))) & (h2(op1(e16, e10)) = op2(h2(e16), h2(e10))) & (h2(op1(e15, e16)) = op2(h2(e15), h2(e16))) & (h2(op1(e15, e15)) = op2(h2(e15), h2(e15))) & (h2(op1(e15, e14)) = op2(h2(e15), h2(e14))) & (h2(op1(e15, e13)) = op2(h2(e15), h2(e13))) & (h2(op1(e15, e12)) = op2(h2(e15), h2(e12))) & (h2(op1(e15, e11)) = op2(h2(e15), h2(e11))) & (h2(op1(e15, e10)) = op2(h2(e15), h2(e10))) & (h2(op1(e14, e16)) = op2(h2(e14), h2(e16))) & (h2(op1(e14, e15)) = op2(h2(e14), h2(e15))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e16)) = op2(h2(e13), h2(e16))) & (h2(op1(e13, e15)) = op2(h2(e13), h2(e15))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e16)) = op2(h2(e12), h2(e16))) & (h2(op1(e12, e15)) = op2(h2(e12), h2(e15))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e16)) = op2(h2(e11), h2(e16))) & (h2(op1(e11, e15)) = op2(h2(e11), h2(e15))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e16)) = op2(h2(e10), h2(e16))) & (h2(op1(e10, e15)) = op2(h2(e10), h2(e15))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e26 = h1(e16)) | (e26 = h1(e15)) | (e26 = h1(e14)) | (e26 = h1(e13)) | (e26 = h1(e12)) | (e26 = h1(e11)) | (e26 = h1(e10))) & ((e25 = h1(e16)) | (e25 = h1(e15)) | (e25 = h1(e14)) | (e25 = h1(e13)) | (e25 = h1(e12)) | (e25 = h1(e11)) | (e25 = h1(e10))) & ((e24 = h1(e16)) | (e24 = h1(e15)) | (e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e16)) | (e23 = h1(e15)) | (e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e16)) | (e22 = h1(e15)) | (e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e16)) | (e21 = h1(e15)) | (e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e16)) | (e20 = h1(e15)) | (e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e16, e16)) = op2(h1(e16), h1(e16))) & (h1(op1(e16, e15)) = op2(h1(e16), h1(e15))) & (h1(op1(e16, e14)) = op2(h1(e16), h1(e14))) & (h1(op1(e16, e13)) = op2(h1(e16), h1(e13))) & (h1(op1(e16, e12)) = op2(h1(e16), h1(e12))) & (h1(op1(e16, e11)) = op2(h1(e16), h1(e11))) & (h1(op1(e16, e10)) = op2(h1(e16), h1(e10))) & (h1(op1(e15, e16)) = op2(h1(e15), h1(e16))) & (h1(op1(e15, e15)) = op2(h1(e15), h1(e15))) & (h1(op1(e15, e14)) = op2(h1(e15), h1(e14))) & (h1(op1(e15, e13)) = op2(h1(e15), h1(e13))) & (h1(op1(e15, e12)) = op2(h1(e15), h1(e12))) & (h1(op1(e15, e11)) = op2(h1(e15), h1(e11))) & (h1(op1(e15, e10)) = op2(h1(e15), h1(e10))) & (h1(op1(e14, e16)) = op2(h1(e14), h1(e16))) & (h1(op1(e14, e15)) = op2(h1(e14), h1(e15))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e16)) = op2(h1(e13), h1(e16))) & (h1(op1(e13, e15)) = op2(h1(e13), h1(e15))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e16)) = op2(h1(e12), h1(e16))) & (h1(op1(e12, e15)) = op2(h1(e12), h1(e15))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e16)) = op2(h1(e11), h1(e16))) & (h1(op1(e11, e15)) = op2(h1(e11), h1(e15))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e16)) = op2(h1(e10), h1(e16))) & (h1(op1(e10, e15)) = op2(h1(e10), h1(e15))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', co1)).
fof(f5200, plain, ~ spl42_749, inference(avatar_split_clause, [], [f5199, f5177])).
fof(f5199, plain, ~ sP30, inference(subsumption_resolution, [], [f1334, f4820])).
fof(f1334, plain, (~ (e20 = h6(e10)) | ~ sP30), inference(cnf_transformation, [], [f80])).
fof(f80, plain, ((~ (e20 = h6(e16)) & ~ (e20 = h6(e15)) & ~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP30), inference(nnf_transformation, [], [f56])).
fof(f5174, plain, (~ spl42_745 | ~ spl42_748), inference(avatar_split_clause, [], [f1328, f5171, f5153])).
fof(f1328, plain, (~ (e21 = h6(e11)) | ~ sP31), inference(cnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e21 = h6(e16)) & ~ (e21 = h6(e15)) & ~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP31), inference(nnf_transformation, [], [f57])).
fof(f5145, plain, (~ spl42_741 | ~ spl42_743), inference(avatar_split_clause, [], [f5140, f5142, f5129])).
fof(f5140, plain, (~ (e22 = h5(e11)) | ~ sP32), inference(forward_demodulation, [], [f1322, f4802])).
fof(f1322, plain, (~ (e22 = h6(e12)) | ~ sP32), inference(cnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e22 = h6(e16)) & ~ (e22 = h6(e15)) & ~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP32), inference(nnf_transformation, [], [f58])).
fof(f5115, plain, ~ spl42_737, inference(avatar_split_clause, [], [f5114, f5104])).
fof(f5114, plain, ~ sP33, inference(subsumption_resolution, [], [f1316, f4800])).
fof(f1316, plain, (~ (e23 = h6(e13)) | ~ sP33), inference(cnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e23 = h6(e16)) & ~ (e23 = h6(e15)) & ~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | ~ sP33), inference(nnf_transformation, [], [f59])).
fof(f5089, plain, ~ spl42_733, inference(avatar_split_clause, [], [f5088, f5079])).
fof(f5088, plain, ~ sP34, inference(subsumption_resolution, [], [f1310, f4797])).
fof(f1310, plain, (~ (e24 = h6(e14)) | ~ sP34), inference(cnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e24 = h6(e16)) & ~ (e24 = h6(e15)) & ~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | ~ sP34), inference(nnf_transformation, [], [f60])).
fof(f5063, plain, ~ spl42_729, inference(avatar_split_clause, [], [f5062, f5054])).
fof(f5062, plain, ~ sP35, inference(subsumption_resolution, [], [f1304, f1243])).
fof(f1304, plain, (~ (e25 = h6(e15)) | ~ sP35), inference(cnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e25 = h6(e16)) & ~ (e25 = h6(e15)) & ~ (e25 = h6(e14)) & ~ (e25 = h6(e13)) & ~ (e25 = h6(e12)) & ~ (e25 = h6(e11)) & ~ (e25 = h6(e10))) | ~ sP35), inference(nnf_transformation, [], [f61])).
fof(f4671, plain, spl42_656, inference(avatar_split_clause, [], [f4670, f4412])).
fof(f4670, plain, (e24 = op2(e20, e24)), inference(backward_demodulation, [], [f486, f4669])).
fof(f486, plain, (e24 = op2(op2(op2(e24, e24), e24), e24)), inference(cnf_transformation, [], [f6])).
fof(f4668, plain, spl42_401, inference(avatar_split_clause, [], [f1203, f3356])).
fof(f1203, plain, (e21 = op2(e25, e25)), inference(cnf_transformation, [], [f15])).
fof(f4667, plain, spl42_458, inference(avatar_split_clause, [], [f4666, f3592])).
fof(f4666, plain, (e22 = op2(e24, e24)), inference(forward_demodulation, [], [f1204, f1206])).
fof(f1204, plain, (e22 = op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25)))), inference(cnf_transformation, [], [f15])).
fof(f4665, plain, spl42_511, inference(avatar_split_clause, [], [f4664, f3811])).
fof(f4664, plain, (e26 = op2(e23, e24)), inference(backward_demodulation, [], [f4662, f4663])).
fof(f4662, plain, (e26 = op2(op2(op2(e24, e24), e25), e24)), inference(backward_demodulation, [], [f1207, f1206])).
fof(f1207, plain, (e26 = op2(op2(op2(op2(e25, op2(e25, e25)), op2(e25, op2(e25, e25))), e25), op2(e25, op2(e25, e25)))), inference(cnf_transformation, [], [f15])).
fof(f4661, plain, spl42_313, inference(avatar_split_clause, [], [f4660, f2893])).
fof(f4660, plain, (e14 = op1(e10, e14)), inference(backward_demodulation, [], [f290, f4659])).
fof(f290, plain, (e14 = op1(op1(op1(e14, e14), e14), e14)), inference(cnf_transformation, [], [f3])).
fof(f4658, plain, spl42_58, inference(avatar_split_clause, [], [f1197, f1837])).
fof(f1197, plain, (e11 = op1(e15, e15)), inference(cnf_transformation, [], [f14])).
fof(f4657, plain, spl42_115, inference(avatar_split_clause, [], [f4656, f2073])).
fof(f4656, plain, (e12 = op1(e14, e14)), inference(forward_demodulation, [], [f1198, f1200])).
fof(f1198, plain, (e12 = op1(op1(e15, op1(e15, e15)), op1(e15, op1(e15, e15)))), inference(cnf_transformation, [], [f14])).
fof(f4655, plain, spl42_168, inference(avatar_split_clause, [], [f4654, f2292])).
fof(f4654, plain, (e16 = op1(e13, e14)), inference(backward_demodulation, [], [f4652, f4653])).
fof(f4652, plain, (e16 = op1(op1(op1(e14, e14), e15), e14)), inference(backward_demodulation, [], [f1201, f1200])).
fof(f1201, plain, (e16 = op1(op1(op1(op1(e15, op1(e15, e15)), op1(e15, op1(e15, e15))), e15), op1(e15, op1(e15, e15)))), inference(cnf_transformation, [], [f14])).
fof(f4651, plain, ~ spl42_680, inference(avatar_split_clause, [], [f1189, f4512])).
fof(f4512, plain, (spl42_680 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl42_680])])).
fof(f1189, plain, ~ (e20 = op2(e20, e20)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, (~ (e26 = op2(e26, e26)) & ~ (e25 = op2(e25, e25)) & ~ (e24 = op2(e24, e24)) & ~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax13)).
fof(f4650, plain, ~ spl42_625, inference(avatar_split_clause, [], [f1190, f4284])).
fof(f4284, plain, (spl42_625 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl42_625])])).
fof(f1190, plain, ~ (e21 = op2(e21, e21)), inference(cnf_transformation, [], [f13])).
fof(f4649, plain, ~ spl42_570, inference(avatar_split_clause, [], [f1191, f4056])).
fof(f4056, plain, (spl42_570 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl42_570])])).
fof(f1191, plain, ~ (e22 = op2(e22, e22)), inference(cnf_transformation, [], [f13])).
fof(f4648, plain, ~ spl42_515, inference(avatar_split_clause, [], [f1192, f3828])).
fof(f3828, plain, (spl42_515 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl42_515])])).
fof(f1192, plain, ~ (e23 = op2(e23, e23)), inference(cnf_transformation, [], [f13])).
fof(f4647, plain, ~ spl42_460, inference(avatar_split_clause, [], [f1193, f3600])).
fof(f3600, plain, (spl42_460 <=> (e24 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl42_460])])).
fof(f1193, plain, ~ (e24 = op2(e24, e24)), inference(cnf_transformation, [], [f13])).
fof(f4645, plain, ~ spl42_350, inference(avatar_split_clause, [], [f1195, f3144])).
fof(f1195, plain, ~ (e26 = op2(e26, e26)), inference(cnf_transformation, [], [f13])).
fof(f4644, plain, ~ spl42_337, inference(avatar_split_clause, [], [f1182, f2993])).
fof(f2993, plain, (spl42_337 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl42_337])])).
fof(f1182, plain, ~ (e10 = op1(e10, e10)), inference(cnf_transformation, [], [f12])).
fof(f12, plain, (~ (e16 = op1(e16, e16)) & ~ (e15 = op1(e15, e15)) & ~ (e14 = op1(e14, e14)) & ~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax12)).
fof(f4643, plain, ~ spl42_282, inference(avatar_split_clause, [], [f1183, f2765])).
fof(f1183, plain, ~ (e11 = op1(e11, e11)), inference(cnf_transformation, [], [f12])).
fof(f4642, plain, ~ spl42_227, inference(avatar_split_clause, [], [f1184, f2537])).
fof(f1184, plain, ~ (e12 = op1(e12, e12)), inference(cnf_transformation, [], [f12])).
fof(f4641, plain, ~ spl42_172, inference(avatar_split_clause, [], [f1185, f2309])).
fof(f2309, plain, (spl42_172 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl42_172])])).
fof(f1185, plain, ~ (e13 = op1(e13, e13)), inference(cnf_transformation, [], [f12])).
fof(f4637, plain, (spl42_680 | spl42_673 | spl42_666 | spl42_659 | spl42_652 | spl42_645 | spl42_638), inference(avatar_split_clause, [], [f356, f4338, f4367, f4396, f4425, f4454, f4483, f4512])).
fof(f356, plain, ((e20 = op2(e20, e26)) | (e20 = op2(e20, e25)) | (e20 = op2(e20, e24)) | (e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e26 = op2(e26, e26)) | (e26 = op2(e25, e26)) | (e26 = op2(e24, e26)) | (e26 = op2(e23, e26)) | (e26 = op2(e22, e26)) | (e26 = op2(e21, e26)) | (e26 = op2(e20, e26))) & ((e26 = op2(e26, e26)) | (e26 = op2(e26, e25)) | (e26 = op2(e26, e24)) | (e26 = op2(e26, e23)) | (e26 = op2(e26, e22)) | (e26 = op2(e26, e21)) | (e26 = op2(e26, e20))) & ((e25 = op2(e26, e26)) | (e25 = op2(e25, e26)) | (e25 = op2(e24, e26)) | (e25 = op2(e23, e26)) | (e25 = op2(e22, e26)) | (e25 = op2(e21, e26)) | (e25 = op2(e20, e26))) & ((e25 = op2(e26, e26)) | (e25 = op2(e26, e25)) | (e25 = op2(e26, e24)) | (e25 = op2(e26, e23)) | (e25 = op2(e26, e22)) | (e25 = op2(e26, e21)) | (e25 = op2(e26, e20))) & ((e24 = op2(e26, e26)) | (e24 = op2(e25, e26)) | (e24 = op2(e24, e26)) | (e24 = op2(e23, e26)) | (e24 = op2(e22, e26)) | (e24 = op2(e21, e26)) | (e24 = op2(e20, e26))) & ((e24 = op2(e26, e26)) | (e24 = op2(e26, e25)) | (e24 = op2(e26, e24)) | (e24 = op2(e26, e23)) | (e24 = op2(e26, e22)) | (e24 = op2(e26, e21)) | (e24 = op2(e26, e20))) & ((e23 = op2(e26, e26)) | (e23 = op2(e25, e26)) | (e23 = op2(e24, e26)) | (e23 = op2(e23, e26)) | (e23 = op2(e22, e26)) | (e23 = op2(e21, e26)) | (e23 = op2(e20, e26))) & ((e23 = op2(e26, e26)) | (e23 = op2(e26, e25)) | (e23 = op2(e26, e24)) | (e23 = op2(e26, e23)) | (e23 = op2(e26, e22)) | (e23 = op2(e26, e21)) | (e23 = op2(e26, e20))) & ((e22 = op2(e26, e26)) | (e22 = op2(e25, e26)) | (e22 = op2(e24, e26)) | (e22 = op2(e23, e26)) | (e22 = op2(e22, e26)) | (e22 = op2(e21, e26)) | (e22 = op2(e20, e26))) & ((e22 = op2(e26, e26)) | (e22 = op2(e26, e25)) | (e22 = op2(e26, e24)) | (e22 = op2(e26, e23)) | (e22 = op2(e26, e22)) | (e22 = op2(e26, e21)) | (e22 = op2(e26, e20))) & ((e21 = op2(e26, e26)) | (e21 = op2(e25, e26)) | (e21 = op2(e24, e26)) | (e21 = op2(e23, e26)) | (e21 = op2(e22, e26)) | (e21 = op2(e21, e26)) | (e21 = op2(e20, e26))) & ((e21 = op2(e26, e26)) | (e21 = op2(e26, e25)) | (e21 = op2(e26, e24)) | (e21 = op2(e26, e23)) | (e21 = op2(e26, e22)) | (e21 = op2(e26, e21)) | (e21 = op2(e26, e20))) & ((e20 = op2(e26, e26)) | (e20 = op2(e25, e26)) | (e20 = op2(e24, e26)) | (e20 = op2(e23, e26)) | (e20 = op2(e22, e26)) | (e20 = op2(e21, e26)) | (e20 = op2(e20, e26))) & ((e20 = op2(e26, e26)) | (e20 = op2(e26, e25)) | (e20 = op2(e26, e24)) | (e20 = op2(e26, e23)) | (e20 = op2(e26, e22)) | (e20 = op2(e26, e21)) | (e20 = op2(e26, e20))) & ((e26 = op2(e26, e25)) | (e26 = op2(e25, e25)) | (e26 = op2(e24, e25)) | (e26 = op2(e23, e25)) | (e26 = op2(e22, e25)) | (e26 = op2(e21, e25)) | (e26 = op2(e20, e25))) & ((e26 = op2(e25, e26)) | (e26 = op2(e25, e25)) | (e26 = op2(e25, e24)) | (e26 = op2(e25, e23)) | (e26 = op2(e25, e22)) | (e26 = op2(e25, e21)) | (e26 = op2(e25, e20))) & ((e25 = op2(e26, e25)) | (e25 = op2(e25, e25)) | (e25 = op2(e24, e25)) | (e25 = op2(e23, e25)) | (e25 = op2(e22, e25)) | (e25 = op2(e21, e25)) | (e25 = op2(e20, e25))) & ((e25 = op2(e25, e26)) | (e25 = op2(e25, e25)) | (e25 = op2(e25, e24)) | (e25 = op2(e25, e23)) | (e25 = op2(e25, e22)) | (e25 = op2(e25, e21)) | (e25 = op2(e25, e20))) & ((e24 = op2(e26, e25)) | (e24 = op2(e25, e25)) | (e24 = op2(e24, e25)) | (e24 = op2(e23, e25)) | (e24 = op2(e22, e25)) | (e24 = op2(e21, e25)) | (e24 = op2(e20, e25))) & ((e24 = op2(e25, e26)) | (e24 = op2(e25, e25)) | (e24 = op2(e25, e24)) | (e24 = op2(e25, e23)) | (e24 = op2(e25, e22)) | (e24 = op2(e25, e21)) | (e24 = op2(e25, e20))) & ((e23 = op2(e26, e25)) | (e23 = op2(e25, e25)) | (e23 = op2(e24, e25)) | (e23 = op2(e23, e25)) | (e23 = op2(e22, e25)) | (e23 = op2(e21, e25)) | (e23 = op2(e20, e25))) & ((e23 = op2(e25, e26)) | (e23 = op2(e25, e25)) | (e23 = op2(e25, e24)) | (e23 = op2(e25, e23)) | (e23 = op2(e25, e22)) | (e23 = op2(e25, e21)) | (e23 = op2(e25, e20))) & ((e22 = op2(e26, e25)) | (e22 = op2(e25, e25)) | (e22 = op2(e24, e25)) | (e22 = op2(e23, e25)) | (e22 = op2(e22, e25)) | (e22 = op2(e21, e25)) | (e22 = op2(e20, e25))) & ((e22 = op2(e25, e26)) | (e22 = op2(e25, e25)) | (e22 = op2(e25, e24)) | (e22 = op2(e25, e23)) | (e22 = op2(e25, e22)) | (e22 = op2(e25, e21)) | (e22 = op2(e25, e20))) & ((e21 = op2(e26, e25)) | (e21 = op2(e25, e25)) | (e21 = op2(e24, e25)) | (e21 = op2(e23, e25)) | (e21 = op2(e22, e25)) | (e21 = op2(e21, e25)) | (e21 = op2(e20, e25))) & ((e21 = op2(e25, e26)) | (e21 = op2(e25, e25)) | (e21 = op2(e25, e24)) | (e21 = op2(e25, e23)) | (e21 = op2(e25, e22)) | (e21 = op2(e25, e21)) | (e21 = op2(e25, e20))) & ((e20 = op2(e26, e25)) | (e20 = op2(e25, e25)) | (e20 = op2(e24, e25)) | (e20 = op2(e23, e25)) | (e20 = op2(e22, e25)) | (e20 = op2(e21, e25)) | (e20 = op2(e20, e25))) & ((e20 = op2(e25, e26)) | (e20 = op2(e25, e25)) | (e20 = op2(e25, e24)) | (e20 = op2(e25, e23)) | (e20 = op2(e25, e22)) | (e20 = op2(e25, e21)) | (e20 = op2(e25, e20))) & ((e26 = op2(e26, e24)) | (e26 = op2(e25, e24)) | (e26 = op2(e24, e24)) | (e26 = op2(e23, e24)) | (e26 = op2(e22, e24)) | (e26 = op2(e21, e24)) | (e26 = op2(e20, e24))) & ((e26 = op2(e24, e26)) | (e26 = op2(e24, e25)) | (e26 = op2(e24, e24)) | (e26 = op2(e24, e23)) | (e26 = op2(e24, e22)) | (e26 = op2(e24, e21)) | (e26 = op2(e24, e20))) & ((e25 = op2(e26, e24)) | (e25 = op2(e25, e24)) | (e25 = op2(e24, e24)) | (e25 = op2(e23, e24)) | (e25 = op2(e22, e24)) | (e25 = op2(e21, e24)) | (e25 = op2(e20, e24))) & ((e25 = op2(e24, e26)) | (e25 = op2(e24, e25)) | (e25 = op2(e24, e24)) | (e25 = op2(e24, e23)) | (e25 = op2(e24, e22)) | (e25 = op2(e24, e21)) | (e25 = op2(e24, e20))) & ((e24 = op2(e26, e24)) | (e24 = op2(e25, e24)) | (e24 = op2(e24, e24)) | (e24 = op2(e23, e24)) | (e24 = op2(e22, e24)) | (e24 = op2(e21, e24)) | (e24 = op2(e20, e24))) & ((e24 = op2(e24, e26)) | (e24 = op2(e24, e25)) | (e24 = op2(e24, e24)) | (e24 = op2(e24, e23)) | (e24 = op2(e24, e22)) | (e24 = op2(e24, e21)) | (e24 = op2(e24, e20))) & ((e23 = op2(e26, e24)) | (e23 = op2(e25, e24)) | (e23 = op2(e24, e24)) | (e23 = op2(e23, e24)) | (e23 = op2(e22, e24)) | (e23 = op2(e21, e24)) | (e23 = op2(e20, e24))) & ((e23 = op2(e24, e26)) | (e23 = op2(e24, e25)) | (e23 = op2(e24, e24)) | (e23 = op2(e24, e23)) | (e23 = op2(e24, e22)) | (e23 = op2(e24, e21)) | (e23 = op2(e24, e20))) & ((e22 = op2(e26, e24)) | (e22 = op2(e25, e24)) | (e22 = op2(e24, e24)) | (e22 = op2(e23, e24)) | (e22 = op2(e22, e24)) | (e22 = op2(e21, e24)) | (e22 = op2(e20, e24))) & ((e22 = op2(e24, e26)) | (e22 = op2(e24, e25)) | (e22 = op2(e24, e24)) | (e22 = op2(e24, e23)) | (e22 = op2(e24, e22)) | (e22 = op2(e24, e21)) | (e22 = op2(e24, e20))) & ((e21 = op2(e26, e24)) | (e21 = op2(e25, e24)) | (e21 = op2(e24, e24)) | (e21 = op2(e23, e24)) | (e21 = op2(e22, e24)) | (e21 = op2(e21, e24)) | (e21 = op2(e20, e24))) & ((e21 = op2(e24, e26)) | (e21 = op2(e24, e25)) | (e21 = op2(e24, e24)) | (e21 = op2(e24, e23)) | (e21 = op2(e24, e22)) | (e21 = op2(e24, e21)) | (e21 = op2(e24, e20))) & ((e20 = op2(e26, e24)) | (e20 = op2(e25, e24)) | (e20 = op2(e24, e24)) | (e20 = op2(e23, e24)) | (e20 = op2(e22, e24)) | (e20 = op2(e21, e24)) | (e20 = op2(e20, e24))) & ((e20 = op2(e24, e26)) | (e20 = op2(e24, e25)) | (e20 = op2(e24, e24)) | (e20 = op2(e24, e23)) | (e20 = op2(e24, e22)) | (e20 = op2(e24, e21)) | (e20 = op2(e24, e20))) & ((e26 = op2(e26, e23)) | (e26 = op2(e25, e23)) | (e26 = op2(e24, e23)) | (e26 = op2(e23, e23)) | (e26 = op2(e22, e23)) | (e26 = op2(e21, e23)) | (e26 = op2(e20, e23))) & ((e26 = op2(e23, e26)) | (e26 = op2(e23, e25)) | (e26 = op2(e23, e24)) | (e26 = op2(e23, e23)) | (e26 = op2(e23, e22)) | (e26 = op2(e23, e21)) | (e26 = op2(e23, e20))) & ((e25 = op2(e26, e23)) | (e25 = op2(e25, e23)) | (e25 = op2(e24, e23)) | (e25 = op2(e23, e23)) | (e25 = op2(e22, e23)) | (e25 = op2(e21, e23)) | (e25 = op2(e20, e23))) & ((e25 = op2(e23, e26)) | (e25 = op2(e23, e25)) | (e25 = op2(e23, e24)) | (e25 = op2(e23, e23)) | (e25 = op2(e23, e22)) | (e25 = op2(e23, e21)) | (e25 = op2(e23, e20))) & ((e24 = op2(e26, e23)) | (e24 = op2(e25, e23)) | (e24 = op2(e24, e23)) | (e24 = op2(e23, e23)) | (e24 = op2(e22, e23)) | (e24 = op2(e21, e23)) | (e24 = op2(e20, e23))) & ((e24 = op2(e23, e26)) | (e24 = op2(e23, e25)) | (e24 = op2(e23, e24)) | (e24 = op2(e23, e23)) | (e24 = op2(e23, e22)) | (e24 = op2(e23, e21)) | (e24 = op2(e23, e20))) & ((e23 = op2(e26, e23)) | (e23 = op2(e25, e23)) | (e23 = op2(e24, e23)) | (e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e26)) | (e23 = op2(e23, e25)) | (e23 = op2(e23, e24)) | (e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e26, e23)) | (e22 = op2(e25, e23)) | (e22 = op2(e24, e23)) | (e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e26)) | (e22 = op2(e23, e25)) | (e22 = op2(e23, e24)) | (e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e26, e23)) | (e21 = op2(e25, e23)) | (e21 = op2(e24, e23)) | (e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e26)) | (e21 = op2(e23, e25)) | (e21 = op2(e23, e24)) | (e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e26, e23)) | (e20 = op2(e25, e23)) | (e20 = op2(e24, e23)) | (e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e26)) | (e20 = op2(e23, e25)) | (e20 = op2(e23, e24)) | (e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e26 = op2(e26, e22)) | (e26 = op2(e25, e22)) | (e26 = op2(e24, e22)) | (e26 = op2(e23, e22)) | (e26 = op2(e22, e22)) | (e26 = op2(e21, e22)) | (e26 = op2(e20, e22))) & ((e26 = op2(e22, e26)) | (e26 = op2(e22, e25)) | (e26 = op2(e22, e24)) | (e26 = op2(e22, e23)) | (e26 = op2(e22, e22)) | (e26 = op2(e22, e21)) | (e26 = op2(e22, e20))) & ((e25 = op2(e26, e22)) | (e25 = op2(e25, e22)) | (e25 = op2(e24, e22)) | (e25 = op2(e23, e22)) | (e25 = op2(e22, e22)) | (e25 = op2(e21, e22)) | (e25 = op2(e20, e22))) & ((e25 = op2(e22, e26)) | (e25 = op2(e22, e25)) | (e25 = op2(e22, e24)) | (e25 = op2(e22, e23)) | (e25 = op2(e22, e22)) | (e25 = op2(e22, e21)) | (e25 = op2(e22, e20))) & ((e24 = op2(e26, e22)) | (e24 = op2(e25, e22)) | (e24 = op2(e24, e22)) | (e24 = op2(e23, e22)) | (e24 = op2(e22, e22)) | (e24 = op2(e21, e22)) | (e24 = op2(e20, e22))) & ((e24 = op2(e22, e26)) | (e24 = op2(e22, e25)) | (e24 = op2(e22, e24)) | (e24 = op2(e22, e23)) | (e24 = op2(e22, e22)) | (e24 = op2(e22, e21)) | (e24 = op2(e22, e20))) & ((e23 = op2(e26, e22)) | (e23 = op2(e25, e22)) | (e23 = op2(e24, e22)) | (e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e26)) | (e23 = op2(e22, e25)) | (e23 = op2(e22, e24)) | (e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e26, e22)) | (e22 = op2(e25, e22)) | (e22 = op2(e24, e22)) | (e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e26)) | (e22 = op2(e22, e25)) | (e22 = op2(e22, e24)) | (e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e26, e22)) | (e21 = op2(e25, e22)) | (e21 = op2(e24, e22)) | (e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e26)) | (e21 = op2(e22, e25)) | (e21 = op2(e22, e24)) | (e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e26, e22)) | (e20 = op2(e25, e22)) | (e20 = op2(e24, e22)) | (e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e26)) | (e20 = op2(e22, e25)) | (e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e26 = op2(e26, e21)) | (e26 = op2(e25, e21)) | (e26 = op2(e24, e21)) | (e26 = op2(e23, e21)) | (e26 = op2(e22, e21)) | (e26 = op2(e21, e21)) | (e26 = op2(e20, e21))) & ((e26 = op2(e21, e26)) | (e26 = op2(e21, e25)) | (e26 = op2(e21, e24)) | (e26 = op2(e21, e23)) | (e26 = op2(e21, e22)) | (e26 = op2(e21, e21)) | (e26 = op2(e21, e20))) & ((e25 = op2(e26, e21)) | (e25 = op2(e25, e21)) | (e25 = op2(e24, e21)) | (e25 = op2(e23, e21)) | (e25 = op2(e22, e21)) | (e25 = op2(e21, e21)) | (e25 = op2(e20, e21))) & ((e25 = op2(e21, e26)) | (e25 = op2(e21, e25)) | (e25 = op2(e21, e24)) | (e25 = op2(e21, e23)) | (e25 = op2(e21, e22)) | (e25 = op2(e21, e21)) | (e25 = op2(e21, e20))) & ((e24 = op2(e26, e21)) | (e24 = op2(e25, e21)) | (e24 = op2(e24, e21)) | (e24 = op2(e23, e21)) | (e24 = op2(e22, e21)) | (e24 = op2(e21, e21)) | (e24 = op2(e20, e21))) & ((e24 = op2(e21, e26)) | (e24 = op2(e21, e25)) | (e24 = op2(e21, e24)) | (e24 = op2(e21, e23)) | (e24 = op2(e21, e22)) | (e24 = op2(e21, e21)) | (e24 = op2(e21, e20))) & ((e23 = op2(e26, e21)) | (e23 = op2(e25, e21)) | (e23 = op2(e24, e21)) | (e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e26)) | (e23 = op2(e21, e25)) | (e23 = op2(e21, e24)) | (e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e26, e21)) | (e22 = op2(e25, e21)) | (e22 = op2(e24, e21)) | (e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e26)) | (e22 = op2(e21, e25)) | (e22 = op2(e21, e24)) | (e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e26, e21)) | (e21 = op2(e25, e21)) | (e21 = op2(e24, e21)) | (e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e26)) | (e21 = op2(e21, e25)) | (e21 = op2(e21, e24)) | (e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e26, e21)) | (e20 = op2(e25, e21)) | (e20 = op2(e24, e21)) | (e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e26)) | (e20 = op2(e21, e25)) | (e20 = op2(e21, e24)) | (e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e26 = op2(e26, e20)) | (e26 = op2(e25, e20)) | (e26 = op2(e24, e20)) | (e26 = op2(e23, e20)) | (e26 = op2(e22, e20)) | (e26 = op2(e21, e20)) | (op2(e20, e20) = e26)) & ((e26 = op2(e20, e26)) | (e26 = op2(e20, e25)) | (e26 = op2(e20, e24)) | (e26 = op2(e20, e23)) | (e26 = op2(e20, e22)) | (e26 = op2(e20, e21)) | (op2(e20, e20) = e26)) & ((e25 = op2(e26, e20)) | (e25 = op2(e25, e20)) | (e25 = op2(e24, e20)) | (e25 = op2(e23, e20)) | (e25 = op2(e22, e20)) | (e25 = op2(e21, e20)) | (op2(e20, e20) = e25)) & ((e25 = op2(e20, e26)) | (e25 = op2(e20, e25)) | (e25 = op2(e20, e24)) | (e25 = op2(e20, e23)) | (e25 = op2(e20, e22)) | (e25 = op2(e20, e21)) | (op2(e20, e20) = e25)) & ((e24 = op2(e26, e20)) | (e24 = op2(e25, e20)) | (e24 = op2(e24, e20)) | (e24 = op2(e23, e20)) | (e24 = op2(e22, e20)) | (e24 = op2(e21, e20)) | (op2(e20, e20) = e24)) & ((e24 = op2(e20, e26)) | (e24 = op2(e20, e25)) | (e24 = op2(e20, e24)) | (e24 = op2(e20, e23)) | (e24 = op2(e20, e22)) | (e24 = op2(e20, e21)) | (op2(e20, e20) = e24)) & ((e23 = op2(e26, e20)) | (e23 = op2(e25, e20)) | (e23 = op2(e24, e20)) | (e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e26)) | (e23 = op2(e20, e25)) | (e23 = op2(e20, e24)) | (e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e26, e20)) | (e22 = op2(e25, e20)) | (e22 = op2(e24, e20)) | (e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e26)) | (e22 = op2(e20, e25)) | (e22 = op2(e20, e24)) | (e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e26, e20)) | (e21 = op2(e25, e20)) | (e21 = op2(e24, e20)) | (e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e26)) | (e21 = op2(e20, e25)) | (e21 = op2(e20, e24)) | (e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e26, e20)) | (e20 = op2(e25, e20)) | (e20 = op2(e24, e20)) | (e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e26)) | (e20 = op2(e20, e25)) | (e20 = op2(e20, e24)) | (e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax5)).
fof(f4636, plain, (spl42_680 | spl42_631 | spl42_582 | spl42_533 | spl42_484 | spl42_435 | spl42_386), inference(avatar_split_clause, [], [f357, f3294, f3497, f3700, f3903, f4106, f4309, f4512])).
fof(f357, plain, ((e20 = op2(e26, e20)) | (e20 = op2(e25, e20)) | (e20 = op2(e24, e20)) | (e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f5])).
fof(f4630, plain, (spl42_683 | spl42_634 | spl42_585 | spl42_536 | spl42_487 | spl42_438 | spl42_389), inference(avatar_split_clause, [], [f363, f3306, f3509, f3712, f3915, f4118, f4321, f4524])).
fof(f363, plain, ((e23 = op2(e26, e20)) | (e23 = op2(e25, e20)) | (e23 = op2(e24, e20)) | (e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)), inference(cnf_transformation, [], [f5])).
fof(f4628, plain, (spl42_684 | spl42_635 | spl42_586 | spl42_537 | spl42_488 | spl42_439 | spl42_390), inference(avatar_split_clause, [], [f365, f3310, f3513, f3716, f3919, f4122, f4325, f4528])).
fof(f365, plain, ((e24 = op2(e26, e20)) | (e24 = op2(e25, e20)) | (e24 = op2(e24, e20)) | (e24 = op2(e23, e20)) | (e24 = op2(e22, e20)) | (e24 = op2(e21, e20)) | (op2(e20, e20) = e24)), inference(cnf_transformation, [], [f5])).
fof(f4624, plain, (spl42_686 | spl42_637 | spl42_588 | spl42_539 | spl42_490 | spl42_441 | spl42_392), inference(avatar_split_clause, [], [f369, f3318, f3521, f3724, f3927, f4130, f4333, f4536])).
fof(f369, plain, ((e26 = op2(e26, e20)) | (e26 = op2(e25, e20)) | (e26 = op2(e24, e20)) | (e26 = op2(e23, e20)) | (e26 = op2(e22, e20)) | (e26 = op2(e21, e20)) | (op2(e20, e20) = e26)), inference(cnf_transformation, [], [f5])).
fof(f4621, plain, (spl42_632 | spl42_625 | spl42_618 | spl42_611 | spl42_604 | spl42_597 | spl42_590), inference(avatar_split_clause, [], [f372, f4139, f4168, f4197, f4226, f4255, f4284, f4313])).
fof(f372, plain, ((e21 = op2(e21, e26)) | (e21 = op2(e21, e25)) | (e21 = op2(e21, e24)) | (e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))), inference(cnf_transformation, [], [f5])).
fof(f4618, plain, (spl42_675 | spl42_626 | spl42_577 | spl42_528 | spl42_479 | spl42_430 | spl42_381), inference(avatar_split_clause, [], [f375, f3273, f3476, f3679, f3882, f4085, f4288, f4491])).
fof(f375, plain, ((e22 = op2(e26, e21)) | (e22 = op2(e25, e21)) | (e22 = op2(e24, e21)) | (e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))), inference(cnf_transformation, [], [f5])).
fof(f4611, plain, (spl42_637 | spl42_630 | spl42_623 | spl42_616 | spl42_609 | spl42_602 | spl42_595), inference(avatar_split_clause, [], [f382, f4159, f4188, f4217, f4246, f4275, f4304, f4333])).
fof(f382, plain, ((e26 = op2(e21, e26)) | (e26 = op2(e21, e25)) | (e26 = op2(e21, e24)) | (e26 = op2(e21, e23)) | (e26 = op2(e21, e22)) | (e26 = op2(e21, e21)) | (e26 = op2(e21, e20))), inference(cnf_transformation, [], [f5])).
fof(f4604, plain, (spl42_668 | spl42_619 | spl42_570 | spl42_521 | spl42_472 | spl42_423 | spl42_374), inference(avatar_split_clause, [], [f389, f3244, f3447, f3650, f3853, f4056, f4259, f4462])).
fof(f389, plain, ((e22 = op2(e26, e22)) | (e22 = op2(e25, e22)) | (e22 = op2(e24, e22)) | (e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))), inference(cnf_transformation, [], [f5])).
fof(f4602, plain, (spl42_669 | spl42_620 | spl42_571 | spl42_522 | spl42_473 | spl42_424 | spl42_375), inference(avatar_split_clause, [], [f391, f3248, f3451, f3654, f3857, f4060, f4263, f4466])).
fof(f391, plain, ((e23 = op2(e26, e22)) | (e23 = op2(e25, e22)) | (e23 = op2(e24, e22)) | (e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))), inference(cnf_transformation, [], [f5])).
fof(f4601, plain, (spl42_586 | spl42_579 | spl42_572 | spl42_565 | spl42_558 | spl42_551 | spl42_544), inference(avatar_split_clause, [], [f392, f3948, f3977, f4006, f4035, f4064, f4093, f4122])).
fof(f392, plain, ((e24 = op2(e22, e26)) | (e24 = op2(e22, e25)) | (e24 = op2(e22, e24)) | (e24 = op2(e22, e23)) | (e24 = op2(e22, e22)) | (e24 = op2(e22, e21)) | (e24 = op2(e22, e20))), inference(cnf_transformation, [], [f5])).
fof(f4600, plain, (spl42_670 | spl42_621 | spl42_572 | spl42_523 | spl42_474 | spl42_425 | spl42_376), inference(avatar_split_clause, [], [f393, f3252, f3455, f3658, f3861, f4064, f4267, f4470])).
fof(f393, plain, ((e24 = op2(e26, e22)) | (e24 = op2(e25, e22)) | (e24 = op2(e24, e22)) | (e24 = op2(e23, e22)) | (e24 = op2(e22, e22)) | (e24 = op2(e21, e22)) | (e24 = op2(e20, e22))), inference(cnf_transformation, [], [f5])).
fof(f4589, plain, (spl42_536 | spl42_529 | spl42_522 | spl42_515 | spl42_508 | spl42_501 | spl42_494), inference(avatar_split_clause, [], [f404, f3741, f3770, f3799, f3828, f3857, f3886, f3915])).
fof(f404, plain, ((e23 = op2(e23, e26)) | (e23 = op2(e23, e25)) | (e23 = op2(e23, e24)) | (e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))), inference(cnf_transformation, [], [f5])).
fof(f4586, plain, (spl42_663 | spl42_614 | spl42_565 | spl42_516 | spl42_467 | spl42_418 | spl42_369), inference(avatar_split_clause, [], [f407, f3223, f3426, f3629, f3832, f4035, f4238, f4441])).
fof(f407, plain, ((e24 = op2(e26, e23)) | (e24 = op2(e25, e23)) | (e24 = op2(e24, e23)) | (e24 = op2(e23, e23)) | (e24 = op2(e22, e23)) | (e24 = op2(e21, e23)) | (e24 = op2(e20, e23))), inference(cnf_transformation, [], [f5])).
fof(f4584, plain, (spl42_664 | spl42_615 | spl42_566 | spl42_517 | spl42_468 | spl42_419 | spl42_370), inference(avatar_split_clause, [], [f409, f3227, f3430, f3633, f3836, f4039, f4242, f4445])).
fof(f409, plain, ((e25 = op2(e26, e23)) | (e25 = op2(e25, e23)) | (e25 = op2(e24, e23)) | (e25 = op2(e23, e23)) | (e25 = op2(e22, e23)) | (e25 = op2(e21, e23)) | (e25 = op2(e20, e23))), inference(cnf_transformation, [], [f5])).
fof(f4581, plain, (spl42_484 | spl42_477 | spl42_470 | spl42_463 | spl42_456 | spl42_449 | spl42_442), inference(avatar_split_clause, [], [f412, f3526, f3555, f3584, f3613, f3642, f3671, f3700])).
fof(f412, plain, ((e20 = op2(e24, e26)) | (e20 = op2(e24, e25)) | (e20 = op2(e24, e24)) | (e20 = op2(e24, e23)) | (e20 = op2(e24, e22)) | (e20 = op2(e24, e21)) | (e20 = op2(e24, e20))), inference(cnf_transformation, [], [f5])).
fof(f4575, plain, (spl42_487 | spl42_480 | spl42_473 | spl42_466 | spl42_459 | spl42_452 | spl42_445), inference(avatar_split_clause, [], [f418, f3538, f3567, f3596, f3625, f3654, f3683, f3712])).
fof(f418, plain, ((e23 = op2(e24, e26)) | (e23 = op2(e24, e25)) | (e23 = op2(e24, e24)) | (e23 = op2(e24, e23)) | (e23 = op2(e24, e22)) | (e23 = op2(e24, e21)) | (e23 = op2(e24, e20))), inference(cnf_transformation, [], [f5])).
fof(f4573, plain, (spl42_488 | spl42_481 | spl42_474 | spl42_467 | spl42_460 | spl42_453 | spl42_446), inference(avatar_split_clause, [], [f420, f3542, f3571, f3600, f3629, f3658, f3687, f3716])).
fof(f420, plain, ((e24 = op2(e24, e26)) | (e24 = op2(e24, e25)) | (e24 = op2(e24, e24)) | (e24 = op2(e24, e23)) | (e24 = op2(e24, e22)) | (e24 = op2(e24, e21)) | (e24 = op2(e24, e20))), inference(cnf_transformation, [], [f5])).
fof(f4567, plain, (spl42_435 | spl42_428 | spl42_421 | spl42_414 | spl42_407 | spl42_400 | spl42_393), inference(avatar_split_clause, [], [f426, f3323, f3352, f3381, f3410, f3439, f3468, f3497])).
fof(f426, plain, ((e20 = op2(e25, e26)) | (e20 = op2(e25, e25)) | (e20 = op2(e25, e24)) | (e20 = op2(e25, e23)) | (e20 = op2(e25, e22)) | (e20 = op2(e25, e21)) | (e20 = op2(e25, e20))), inference(cnf_transformation, [], [f5])).
fof(f4563, plain, (spl42_437 | spl42_430 | spl42_423 | spl42_416 | spl42_409 | spl42_402 | spl42_395), inference(avatar_split_clause, [], [f430, f3331, f3360, f3389, f3418, f3447, f3476, f3505])).
fof(f430, plain, ((e22 = op2(e25, e26)) | (e22 = op2(e25, e25)) | (e22 = op2(e25, e24)) | (e22 = op2(e25, e23)) | (e22 = op2(e25, e22)) | (e22 = op2(e25, e21)) | (e22 = op2(e25, e20))), inference(cnf_transformation, [], [f5])).
fof(f4552, plain, (spl42_638 | spl42_589 | spl42_540 | spl42_491 | spl42_442 | spl42_393 | spl42_344), inference(avatar_split_clause, [], [f441, f3120, f3323, f3526, f3729, f3932, f4135, f4338])).
fof(f441, plain, ((e20 = op2(e26, e26)) | (e20 = op2(e25, e26)) | (e20 = op2(e24, e26)) | (e20 = op2(e23, e26)) | (e20 = op2(e22, e26)) | (e20 = op2(e21, e26)) | (e20 = op2(e20, e26))), inference(cnf_transformation, [], [f5])).
fof(f4543, plain, (spl42_391 | spl42_384 | spl42_377 | spl42_370 | spl42_363 | spl42_356 | spl42_349), inference(avatar_split_clause, [], [f450, f3140, f3169, f3198, f3227, f3256, f3285, f3314])).
fof(f450, plain, ((e25 = op2(e26, e26)) | (e25 = op2(e26, e25)) | (e25 = op2(e26, e24)) | (e25 = op2(e26, e23)) | (e25 = op2(e26, e22)) | (e25 = op2(e26, e21)) | (e25 = op2(e26, e20))), inference(cnf_transformation, [], [f5])).
fof(f4542, plain, (spl42_643 | spl42_594 | spl42_545 | spl42_496 | spl42_447 | spl42_398 | spl42_349), inference(avatar_split_clause, [], [f451, f3140, f3343, f3546, f3749, f3952, f4155, f4358])).
fof(f451, plain, ((e25 = op2(e26, e26)) | (e25 = op2(e25, e26)) | (e25 = op2(e24, e26)) | (e25 = op2(e23, e26)) | (e25 = op2(e22, e26)) | (e25 = op2(e21, e26)) | (e25 = op2(e20, e26))), inference(cnf_transformation, [], [f5])).
fof(f4541, plain, (spl42_392 | spl42_385 | spl42_378 | spl42_371 | spl42_364 | spl42_357 | spl42_350), inference(avatar_split_clause, [], [f452, f3144, f3173, f3202, f3231, f3260, f3289, f3318])).
fof(f452, plain, ((e26 = op2(e26, e26)) | (e26 = op2(e26, e25)) | (e26 = op2(e26, e24)) | (e26 = op2(e26, e23)) | (e26 = op2(e26, e22)) | (e26 = op2(e26, e21)) | (e26 = op2(e26, e20))), inference(cnf_transformation, [], [f5])).
fof(f4220, plain, (spl42_603 | spl42_604 | spl42_605 | spl42_606 | spl42_607 | spl42_608 | spl42_609), inference(avatar_split_clause, [], [f318, f4217, f4213, f4209, f4205, f4201, f4197, f4193])).
fof(f318, plain, ((e26 = op2(e21, e24)) | (e25 = op2(e21, e24)) | (e24 = op2(e21, e24)) | (e23 = op2(e21, e24)) | (e22 = op2(e21, e24)) | (e21 = op2(e21, e24)) | (e20 = op2(e21, e24))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e26 = op2(e26, e26)) | (e25 = op2(e26, e26)) | (e24 = op2(e26, e26)) | (e23 = op2(e26, e26)) | (e22 = op2(e26, e26)) | (e21 = op2(e26, e26)) | (e20 = op2(e26, e26))) & ((e26 = op2(e26, e25)) | (e25 = op2(e26, e25)) | (e24 = op2(e26, e25)) | (e23 = op2(e26, e25)) | (e22 = op2(e26, e25)) | (e21 = op2(e26, e25)) | (e20 = op2(e26, e25))) & ((e26 = op2(e26, e24)) | (e25 = op2(e26, e24)) | (e24 = op2(e26, e24)) | (e23 = op2(e26, e24)) | (e22 = op2(e26, e24)) | (e21 = op2(e26, e24)) | (e20 = op2(e26, e24))) & ((e26 = op2(e26, e23)) | (e25 = op2(e26, e23)) | (e24 = op2(e26, e23)) | (e23 = op2(e26, e23)) | (e22 = op2(e26, e23)) | (e21 = op2(e26, e23)) | (e20 = op2(e26, e23))) & ((e26 = op2(e26, e22)) | (e25 = op2(e26, e22)) | (e24 = op2(e26, e22)) | (e23 = op2(e26, e22)) | (e22 = op2(e26, e22)) | (e21 = op2(e26, e22)) | (e20 = op2(e26, e22))) & ((e26 = op2(e26, e21)) | (e25 = op2(e26, e21)) | (e24 = op2(e26, e21)) | (e23 = op2(e26, e21)) | (e22 = op2(e26, e21)) | (e21 = op2(e26, e21)) | (e20 = op2(e26, e21))) & ((e26 = op2(e26, e20)) | (e25 = op2(e26, e20)) | (e24 = op2(e26, e20)) | (e23 = op2(e26, e20)) | (e22 = op2(e26, e20)) | (e21 = op2(e26, e20)) | (e20 = op2(e26, e20))) & ((e26 = op2(e25, e26)) | (e25 = op2(e25, e26)) | (e24 = op2(e25, e26)) | (e23 = op2(e25, e26)) | (e22 = op2(e25, e26)) | (e21 = op2(e25, e26)) | (e20 = op2(e25, e26))) & ((e26 = op2(e25, e25)) | (e25 = op2(e25, e25)) | (e24 = op2(e25, e25)) | (e23 = op2(e25, e25)) | (e22 = op2(e25, e25)) | (e21 = op2(e25, e25)) | (e20 = op2(e25, e25))) & ((e26 = op2(e25, e24)) | (e25 = op2(e25, e24)) | (e24 = op2(e25, e24)) | (e23 = op2(e25, e24)) | (e22 = op2(e25, e24)) | (e21 = op2(e25, e24)) | (e20 = op2(e25, e24))) & ((e26 = op2(e25, e23)) | (e25 = op2(e25, e23)) | (e24 = op2(e25, e23)) | (e23 = op2(e25, e23)) | (e22 = op2(e25, e23)) | (e21 = op2(e25, e23)) | (e20 = op2(e25, e23))) & ((e26 = op2(e25, e22)) | (e25 = op2(e25, e22)) | (e24 = op2(e25, e22)) | (e23 = op2(e25, e22)) | (e22 = op2(e25, e22)) | (e21 = op2(e25, e22)) | (e20 = op2(e25, e22))) & ((e26 = op2(e25, e21)) | (e25 = op2(e25, e21)) | (e24 = op2(e25, e21)) | (e23 = op2(e25, e21)) | (e22 = op2(e25, e21)) | (e21 = op2(e25, e21)) | (e20 = op2(e25, e21))) & ((e26 = op2(e25, e20)) | (e25 = op2(e25, e20)) | (e24 = op2(e25, e20)) | (e23 = op2(e25, e20)) | (e22 = op2(e25, e20)) | (e21 = op2(e25, e20)) | (e20 = op2(e25, e20))) & ((e26 = op2(e24, e26)) | (e25 = op2(e24, e26)) | (e24 = op2(e24, e26)) | (e23 = op2(e24, e26)) | (e22 = op2(e24, e26)) | (e21 = op2(e24, e26)) | (e20 = op2(e24, e26))) & ((e26 = op2(e24, e25)) | (e25 = op2(e24, e25)) | (e24 = op2(e24, e25)) | (e23 = op2(e24, e25)) | (e22 = op2(e24, e25)) | (e21 = op2(e24, e25)) | (e20 = op2(e24, e25))) & ((e26 = op2(e24, e24)) | (e25 = op2(e24, e24)) | (e24 = op2(e24, e24)) | (e23 = op2(e24, e24)) | (e22 = op2(e24, e24)) | (e21 = op2(e24, e24)) | (e20 = op2(e24, e24))) & ((e26 = op2(e24, e23)) | (e25 = op2(e24, e23)) | (e24 = op2(e24, e23)) | (e23 = op2(e24, e23)) | (e22 = op2(e24, e23)) | (e21 = op2(e24, e23)) | (e20 = op2(e24, e23))) & ((e26 = op2(e24, e22)) | (e25 = op2(e24, e22)) | (e24 = op2(e24, e22)) | (e23 = op2(e24, e22)) | (e22 = op2(e24, e22)) | (e21 = op2(e24, e22)) | (e20 = op2(e24, e22))) & ((e26 = op2(e24, e21)) | (e25 = op2(e24, e21)) | (e24 = op2(e24, e21)) | (e23 = op2(e24, e21)) | (e22 = op2(e24, e21)) | (e21 = op2(e24, e21)) | (e20 = op2(e24, e21))) & ((e26 = op2(e24, e20)) | (e25 = op2(e24, e20)) | (e24 = op2(e24, e20)) | (e23 = op2(e24, e20)) | (e22 = op2(e24, e20)) | (e21 = op2(e24, e20)) | (e20 = op2(e24, e20))) & ((e26 = op2(e23, e26)) | (e25 = op2(e23, e26)) | (e24 = op2(e23, e26)) | (e23 = op2(e23, e26)) | (e22 = op2(e23, e26)) | (e21 = op2(e23, e26)) | (e20 = op2(e23, e26))) & ((e26 = op2(e23, e25)) | (e25 = op2(e23, e25)) | (e24 = op2(e23, e25)) | (e23 = op2(e23, e25)) | (e22 = op2(e23, e25)) | (e21 = op2(e23, e25)) | (e20 = op2(e23, e25))) & ((e26 = op2(e23, e24)) | (e25 = op2(e23, e24)) | (e24 = op2(e23, e24)) | (e23 = op2(e23, e24)) | (e22 = op2(e23, e24)) | (e21 = op2(e23, e24)) | (e20 = op2(e23, e24))) & ((e26 = op2(e23, e23)) | (e25 = op2(e23, e23)) | (e24 = op2(e23, e23)) | (e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e26 = op2(e23, e22)) | (e25 = op2(e23, e22)) | (e24 = op2(e23, e22)) | (e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e26 = op2(e23, e21)) | (e25 = op2(e23, e21)) | (e24 = op2(e23, e21)) | (e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e26 = op2(e23, e20)) | (e25 = op2(e23, e20)) | (e24 = op2(e23, e20)) | (e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e26 = op2(e22, e26)) | (e25 = op2(e22, e26)) | (e24 = op2(e22, e26)) | (e23 = op2(e22, e26)) | (e22 = op2(e22, e26)) | (e21 = op2(e22, e26)) | (e20 = op2(e22, e26))) & ((e26 = op2(e22, e25)) | (e25 = op2(e22, e25)) | (e24 = op2(e22, e25)) | (e23 = op2(e22, e25)) | (e22 = op2(e22, e25)) | (e21 = op2(e22, e25)) | (e20 = op2(e22, e25))) & ((e26 = op2(e22, e24)) | (e25 = op2(e22, e24)) | (e24 = op2(e22, e24)) | (e23 = op2(e22, e24)) | (e22 = op2(e22, e24)) | (e21 = op2(e22, e24)) | (e20 = op2(e22, e24))) & ((e26 = op2(e22, e23)) | (e25 = op2(e22, e23)) | (e24 = op2(e22, e23)) | (e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e26 = op2(e22, e22)) | (e25 = op2(e22, e22)) | (e24 = op2(e22, e22)) | (e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e26 = op2(e22, e21)) | (e25 = op2(e22, e21)) | (e24 = op2(e22, e21)) | (e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e26 = op2(e22, e20)) | (e25 = op2(e22, e20)) | (e24 = op2(e22, e20)) | (e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e26 = op2(e21, e26)) | (e25 = op2(e21, e26)) | (e24 = op2(e21, e26)) | (e23 = op2(e21, e26)) | (e22 = op2(e21, e26)) | (e21 = op2(e21, e26)) | (e20 = op2(e21, e26))) & ((e26 = op2(e21, e25)) | (e25 = op2(e21, e25)) | (e24 = op2(e21, e25)) | (e23 = op2(e21, e25)) | (e22 = op2(e21, e25)) | (e21 = op2(e21, e25)) | (e20 = op2(e21, e25))) & ((e26 = op2(e21, e24)) | (e25 = op2(e21, e24)) | (e24 = op2(e21, e24)) | (e23 = op2(e21, e24)) | (e22 = op2(e21, e24)) | (e21 = op2(e21, e24)) | (e20 = op2(e21, e24))) & ((e26 = op2(e21, e23)) | (e25 = op2(e21, e23)) | (e24 = op2(e21, e23)) | (e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e26 = op2(e21, e22)) | (e25 = op2(e21, e22)) | (e24 = op2(e21, e22)) | (e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e26 = op2(e21, e21)) | (e25 = op2(e21, e21)) | (e24 = op2(e21, e21)) | (e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e26 = op2(e21, e20)) | (e25 = op2(e21, e20)) | (e24 = op2(e21, e20)) | (e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e26 = op2(e20, e26)) | (e25 = op2(e20, e26)) | (e24 = op2(e20, e26)) | (e23 = op2(e20, e26)) | (e22 = op2(e20, e26)) | (e21 = op2(e20, e26)) | (e20 = op2(e20, e26))) & ((e26 = op2(e20, e25)) | (e25 = op2(e20, e25)) | (e24 = op2(e20, e25)) | (e23 = op2(e20, e25)) | (e22 = op2(e20, e25)) | (e21 = op2(e20, e25)) | (e20 = op2(e20, e25))) & ((e26 = op2(e20, e24)) | (e25 = op2(e20, e24)) | (e24 = op2(e20, e24)) | (e23 = op2(e20, e24)) | (e22 = op2(e20, e24)) | (e21 = op2(e20, e24)) | (e20 = op2(e20, e24))) & ((e26 = op2(e20, e23)) | (e25 = op2(e20, e23)) | (e24 = op2(e20, e23)) | (e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e26 = op2(e20, e22)) | (e25 = op2(e20, e22)) | (e24 = op2(e20, e22)) | (e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e26 = op2(e20, e21)) | (e25 = op2(e20, e21)) | (e24 = op2(e20, e21)) | (e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e26) | (op2(e20, e20) = e25) | (op2(e20, e20) = e24) | (op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax4)).
fof(f4162, plain, (spl42_589 | spl42_590 | spl42_591 | spl42_592 | spl42_593 | spl42_594 | spl42_595), inference(avatar_split_clause, [], [f320, f4159, f4155, f4151, f4147, f4143, f4139, f4135])).
fof(f320, plain, ((e26 = op2(e21, e26)) | (e25 = op2(e21, e26)) | (e24 = op2(e21, e26)) | (e23 = op2(e21, e26)) | (e22 = op2(e21, e26)) | (e21 = op2(e21, e26)) | (e20 = op2(e21, e26))), inference(cnf_transformation, [], [f4])).
fof(f4133, plain, (spl42_582 | spl42_583 | spl42_584 | spl42_585 | spl42_586 | spl42_587 | spl42_588), inference(avatar_split_clause, [], [f321, f4130, f4126, f4122, f4118, f4114, f4110, f4106])).
fof(f321, plain, ((e26 = op2(e22, e20)) | (e25 = op2(e22, e20)) | (e24 = op2(e22, e20)) | (e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f4104, plain, (spl42_575 | spl42_576 | spl42_577 | spl42_578 | spl42_579 | spl42_580 | spl42_581), inference(avatar_split_clause, [], [f322, f4101, f4097, f4093, f4089, f4085, f4081, f4077])).
fof(f322, plain, ((e26 = op2(e22, e21)) | (e25 = op2(e22, e21)) | (e24 = op2(e22, e21)) | (e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))), inference(cnf_transformation, [], [f4])).
fof(f4046, plain, (spl42_561 | spl42_562 | spl42_563 | spl42_564 | spl42_565 | spl42_566 | spl42_567), inference(avatar_split_clause, [], [f324, f4043, f4039, f4035, f4031, f4027, f4023, f4019])).
fof(f324, plain, ((e26 = op2(e22, e23)) | (e25 = op2(e22, e23)) | (e24 = op2(e22, e23)) | (e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))), inference(cnf_transformation, [], [f4])).
fof(f3640, plain, (spl42_463 | spl42_464 | spl42_465 | spl42_466 | spl42_467 | spl42_468 | spl42_469), inference(avatar_split_clause, [], [f338, f3637, f3633, f3629, f3625, f3621, f3617, f3613])).
fof(f338, plain, ((e26 = op2(e24, e23)) | (e25 = op2(e24, e23)) | (e24 = op2(e24, e23)) | (e23 = op2(e24, e23)) | (e22 = op2(e24, e23)) | (e21 = op2(e24, e23)) | (e20 = op2(e24, e23))), inference(cnf_transformation, [], [f4])).
fof(f3582, plain, (spl42_449 | spl42_450 | spl42_451 | spl42_452 | spl42_453 | spl42_454 | spl42_455), inference(avatar_split_clause, [], [f340, f3579, f3575, f3571, f3567, f3563, f3559, f3555])).
fof(f340, plain, ((e26 = op2(e24, e25)) | (e25 = op2(e24, e25)) | (e24 = op2(e24, e25)) | (e23 = op2(e24, e25)) | (e22 = op2(e24, e25)) | (e21 = op2(e24, e25)) | (e20 = op2(e24, e25))), inference(cnf_transformation, [], [f4])).
fof(f3553, plain, (spl42_442 | spl42_443 | spl42_444 | spl42_445 | spl42_446 | spl42_447 | spl42_448), inference(avatar_split_clause, [], [f341, f3550, f3546, f3542, f3538, f3534, f3530, f3526])).
fof(f341, plain, ((e26 = op2(e24, e26)) | (e25 = op2(e24, e26)) | (e24 = op2(e24, e26)) | (e23 = op2(e24, e26)) | (e22 = op2(e24, e26)) | (e21 = op2(e24, e26)) | (e20 = op2(e24, e26))), inference(cnf_transformation, [], [f4])).
fof(f3408, plain, (spl42_407 | spl42_408 | spl42_409 | spl42_410 | spl42_411 | spl42_412 | spl42_413), inference(avatar_split_clause, [], [f346, f3405, f3401, f3397, f3393, f3389, f3385, f3381])).
fof(f346, plain, ((e26 = op2(e25, e24)) | (e25 = op2(e25, e24)) | (e24 = op2(e25, e24)) | (e23 = op2(e25, e24)) | (e22 = op2(e25, e24)) | (e21 = op2(e25, e24)) | (e20 = op2(e25, e24))), inference(cnf_transformation, [], [f4])).
fof(f3350, plain, (spl42_393 | spl42_394 | spl42_395 | spl42_396 | spl42_397 | spl42_398 | spl42_399), inference(avatar_split_clause, [], [f348, f3347, f3343, f3339, f3335, f3331, f3327, f3323])).
fof(f348, plain, ((e26 = op2(e25, e26)) | (e25 = op2(e25, e26)) | (e24 = op2(e25, e26)) | (e23 = op2(e25, e26)) | (e22 = op2(e25, e26)) | (e21 = op2(e25, e26)) | (e20 = op2(e25, e26))), inference(cnf_transformation, [], [f4])).
fof(f3205, plain, (spl42_358 | spl42_359 | spl42_360 | spl42_361 | spl42_362 | spl42_363 | spl42_364), inference(avatar_split_clause, [], [f353, f3202, f3198, f3194, f3190, f3186, f3182, f3178])).
fof(f353, plain, ((e26 = op2(e26, e24)) | (e25 = op2(e26, e24)) | (e24 = op2(e26, e24)) | (e23 = op2(e26, e24)) | (e22 = op2(e26, e24)) | (e21 = op2(e26, e24)) | (e20 = op2(e26, e24))), inference(cnf_transformation, [], [f4])).
fof(f3118, plain, (spl42_337 | spl42_330 | spl42_323 | spl42_316 | spl42_309 | spl42_302 | spl42_295), inference(avatar_split_clause, [], [f160, f2819, f2848, f2877, f2906, f2935, f2964, f2993])).
fof(f160, plain, ((e10 = op1(e10, e16)) | (e10 = op1(e10, e15)) | (e10 = op1(e10, e14)) | (e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e16 = op1(e16, e16)) | (e16 = op1(e15, e16)) | (e16 = op1(e14, e16)) | (e16 = op1(e13, e16)) | (e16 = op1(e12, e16)) | (e16 = op1(e11, e16)) | (e16 = op1(e10, e16))) & ((e16 = op1(e16, e16)) | (e16 = op1(e16, e15)) | (e16 = op1(e16, e14)) | (e16 = op1(e16, e13)) | (e16 = op1(e16, e12)) | (e16 = op1(e16, e11)) | (e16 = op1(e16, e10))) & ((e15 = op1(e16, e16)) | (e15 = op1(e15, e16)) | (e15 = op1(e14, e16)) | (e15 = op1(e13, e16)) | (e15 = op1(e12, e16)) | (e15 = op1(e11, e16)) | (e15 = op1(e10, e16))) & ((e15 = op1(e16, e16)) | (e15 = op1(e16, e15)) | (e15 = op1(e16, e14)) | (e15 = op1(e16, e13)) | (e15 = op1(e16, e12)) | (e15 = op1(e16, e11)) | (e15 = op1(e16, e10))) & ((e14 = op1(e16, e16)) | (e14 = op1(e15, e16)) | (e14 = op1(e14, e16)) | (e14 = op1(e13, e16)) | (e14 = op1(e12, e16)) | (e14 = op1(e11, e16)) | (e14 = op1(e10, e16))) & ((e14 = op1(e16, e16)) | (e14 = op1(e16, e15)) | (e14 = op1(e16, e14)) | (e14 = op1(e16, e13)) | (e14 = op1(e16, e12)) | (e14 = op1(e16, e11)) | (e14 = op1(e16, e10))) & ((e13 = op1(e16, e16)) | (e13 = op1(e15, e16)) | (e13 = op1(e14, e16)) | (e13 = op1(e13, e16)) | (e13 = op1(e12, e16)) | (e13 = op1(e11, e16)) | (e13 = op1(e10, e16))) & ((e13 = op1(e16, e16)) | (e13 = op1(e16, e15)) | (e13 = op1(e16, e14)) | (e13 = op1(e16, e13)) | (e13 = op1(e16, e12)) | (e13 = op1(e16, e11)) | (e13 = op1(e16, e10))) & ((e12 = op1(e16, e16)) | (e12 = op1(e15, e16)) | (e12 = op1(e14, e16)) | (e12 = op1(e13, e16)) | (e12 = op1(e12, e16)) | (e12 = op1(e11, e16)) | (e12 = op1(e10, e16))) & ((e12 = op1(e16, e16)) | (e12 = op1(e16, e15)) | (e12 = op1(e16, e14)) | (e12 = op1(e16, e13)) | (e12 = op1(e16, e12)) | (e12 = op1(e16, e11)) | (e12 = op1(e16, e10))) & ((e11 = op1(e16, e16)) | (e11 = op1(e15, e16)) | (e11 = op1(e14, e16)) | (e11 = op1(e13, e16)) | (e11 = op1(e12, e16)) | (e11 = op1(e11, e16)) | (e11 = op1(e10, e16))) & ((e11 = op1(e16, e16)) | (e11 = op1(e16, e15)) | (e11 = op1(e16, e14)) | (e11 = op1(e16, e13)) | (e11 = op1(e16, e12)) | (e11 = op1(e16, e11)) | (e11 = op1(e16, e10))) & ((e10 = op1(e16, e16)) | (e10 = op1(e15, e16)) | (e10 = op1(e14, e16)) | (e10 = op1(e13, e16)) | (e10 = op1(e12, e16)) | (e10 = op1(e11, e16)) | (e10 = op1(e10, e16))) & ((e10 = op1(e16, e16)) | (e10 = op1(e16, e15)) | (e10 = op1(e16, e14)) | (e10 = op1(e16, e13)) | (e10 = op1(e16, e12)) | (e10 = op1(e16, e11)) | (e10 = op1(e16, e10))) & ((e16 = op1(e16, e15)) | (e16 = op1(e15, e15)) | (e16 = op1(e14, e15)) | (e16 = op1(e13, e15)) | (e16 = op1(e12, e15)) | (e16 = op1(e11, e15)) | (e16 = op1(e10, e15))) & ((e16 = op1(e15, e16)) | (e16 = op1(e15, e15)) | (e16 = op1(e15, e14)) | (e16 = op1(e15, e13)) | (e16 = op1(e15, e12)) | (e16 = op1(e15, e11)) | (e16 = op1(e15, e10))) & ((e15 = op1(e16, e15)) | (e15 = op1(e15, e15)) | (e15 = op1(e14, e15)) | (e15 = op1(e13, e15)) | (e15 = op1(e12, e15)) | (e15 = op1(e11, e15)) | (e15 = op1(e10, e15))) & ((e15 = op1(e15, e16)) | (e15 = op1(e15, e15)) | (e15 = op1(e15, e14)) | (e15 = op1(e15, e13)) | (e15 = op1(e15, e12)) | (e15 = op1(e15, e11)) | (e15 = op1(e15, e10))) & ((e14 = op1(e16, e15)) | (e14 = op1(e15, e15)) | (e14 = op1(e14, e15)) | (e14 = op1(e13, e15)) | (e14 = op1(e12, e15)) | (e14 = op1(e11, e15)) | (e14 = op1(e10, e15))) & ((e14 = op1(e15, e16)) | (e14 = op1(e15, e15)) | (e14 = op1(e15, e14)) | (e14 = op1(e15, e13)) | (e14 = op1(e15, e12)) | (e14 = op1(e15, e11)) | (e14 = op1(e15, e10))) & ((e13 = op1(e16, e15)) | (e13 = op1(e15, e15)) | (e13 = op1(e14, e15)) | (e13 = op1(e13, e15)) | (e13 = op1(e12, e15)) | (e13 = op1(e11, e15)) | (e13 = op1(e10, e15))) & ((e13 = op1(e15, e16)) | (e13 = op1(e15, e15)) | (e13 = op1(e15, e14)) | (e13 = op1(e15, e13)) | (e13 = op1(e15, e12)) | (e13 = op1(e15, e11)) | (e13 = op1(e15, e10))) & ((e12 = op1(e16, e15)) | (e12 = op1(e15, e15)) | (e12 = op1(e14, e15)) | (e12 = op1(e13, e15)) | (e12 = op1(e12, e15)) | (e12 = op1(e11, e15)) | (e12 = op1(e10, e15))) & ((e12 = op1(e15, e16)) | (e12 = op1(e15, e15)) | (e12 = op1(e15, e14)) | (e12 = op1(e15, e13)) | (e12 = op1(e15, e12)) | (e12 = op1(e15, e11)) | (e12 = op1(e15, e10))) & ((e11 = op1(e16, e15)) | (e11 = op1(e15, e15)) | (e11 = op1(e14, e15)) | (e11 = op1(e13, e15)) | (e11 = op1(e12, e15)) | (e11 = op1(e11, e15)) | (e11 = op1(e10, e15))) & ((e11 = op1(e15, e16)) | (e11 = op1(e15, e15)) | (e11 = op1(e15, e14)) | (e11 = op1(e15, e13)) | (e11 = op1(e15, e12)) | (e11 = op1(e15, e11)) | (e11 = op1(e15, e10))) & ((e10 = op1(e16, e15)) | (e10 = op1(e15, e15)) | (e10 = op1(e14, e15)) | (e10 = op1(e13, e15)) | (e10 = op1(e12, e15)) | (e10 = op1(e11, e15)) | (e10 = op1(e10, e15))) & ((e10 = op1(e15, e16)) | (e10 = op1(e15, e15)) | (e10 = op1(e15, e14)) | (e10 = op1(e15, e13)) | (e10 = op1(e15, e12)) | (e10 = op1(e15, e11)) | (e10 = op1(e15, e10))) & ((e16 = op1(e16, e14)) | (e16 = op1(e15, e14)) | (e16 = op1(e14, e14)) | (e16 = op1(e13, e14)) | (e16 = op1(e12, e14)) | (e16 = op1(e11, e14)) | (e16 = op1(e10, e14))) & ((e16 = op1(e14, e16)) | (e16 = op1(e14, e15)) | (e16 = op1(e14, e14)) | (e16 = op1(e14, e13)) | (e16 = op1(e14, e12)) | (e16 = op1(e14, e11)) | (e16 = op1(e14, e10))) & ((e15 = op1(e16, e14)) | (e15 = op1(e15, e14)) | (e15 = op1(e14, e14)) | (e15 = op1(e13, e14)) | (e15 = op1(e12, e14)) | (e15 = op1(e11, e14)) | (e15 = op1(e10, e14))) & ((e15 = op1(e14, e16)) | (e15 = op1(e14, e15)) | (e15 = op1(e14, e14)) | (e15 = op1(e14, e13)) | (e15 = op1(e14, e12)) | (e15 = op1(e14, e11)) | (e15 = op1(e14, e10))) & ((e14 = op1(e16, e14)) | (e14 = op1(e15, e14)) | (e14 = op1(e14, e14)) | (e14 = op1(e13, e14)) | (e14 = op1(e12, e14)) | (e14 = op1(e11, e14)) | (e14 = op1(e10, e14))) & ((e14 = op1(e14, e16)) | (e14 = op1(e14, e15)) | (e14 = op1(e14, e14)) | (e14 = op1(e14, e13)) | (e14 = op1(e14, e12)) | (e14 = op1(e14, e11)) | (e14 = op1(e14, e10))) & ((e13 = op1(e16, e14)) | (e13 = op1(e15, e14)) | (e13 = op1(e14, e14)) | (e13 = op1(e13, e14)) | (e13 = op1(e12, e14)) | (e13 = op1(e11, e14)) | (e13 = op1(e10, e14))) & ((e13 = op1(e14, e16)) | (e13 = op1(e14, e15)) | (e13 = op1(e14, e14)) | (e13 = op1(e14, e13)) | (e13 = op1(e14, e12)) | (e13 = op1(e14, e11)) | (e13 = op1(e14, e10))) & ((e12 = op1(e16, e14)) | (e12 = op1(e15, e14)) | (e12 = op1(e14, e14)) | (e12 = op1(e13, e14)) | (e12 = op1(e12, e14)) | (e12 = op1(e11, e14)) | (e12 = op1(e10, e14))) & ((e12 = op1(e14, e16)) | (e12 = op1(e14, e15)) | (e12 = op1(e14, e14)) | (e12 = op1(e14, e13)) | (e12 = op1(e14, e12)) | (e12 = op1(e14, e11)) | (e12 = op1(e14, e10))) & ((e11 = op1(e16, e14)) | (e11 = op1(e15, e14)) | (e11 = op1(e14, e14)) | (e11 = op1(e13, e14)) | (e11 = op1(e12, e14)) | (e11 = op1(e11, e14)) | (e11 = op1(e10, e14))) & ((e11 = op1(e14, e16)) | (e11 = op1(e14, e15)) | (e11 = op1(e14, e14)) | (e11 = op1(e14, e13)) | (e11 = op1(e14, e12)) | (e11 = op1(e14, e11)) | (e11 = op1(e14, e10))) & ((e10 = op1(e16, e14)) | (e10 = op1(e15, e14)) | (e10 = op1(e14, e14)) | (e10 = op1(e13, e14)) | (e10 = op1(e12, e14)) | (e10 = op1(e11, e14)) | (e10 = op1(e10, e14))) & ((e10 = op1(e14, e16)) | (e10 = op1(e14, e15)) | (e10 = op1(e14, e14)) | (e10 = op1(e14, e13)) | (e10 = op1(e14, e12)) | (e10 = op1(e14, e11)) | (e10 = op1(e14, e10))) & ((e16 = op1(e16, e13)) | (e16 = op1(e15, e13)) | (e16 = op1(e14, e13)) | (e16 = op1(e13, e13)) | (e16 = op1(e12, e13)) | (e16 = op1(e11, e13)) | (e16 = op1(e10, e13))) & ((e16 = op1(e13, e16)) | (e16 = op1(e13, e15)) | (e16 = op1(e13, e14)) | (e16 = op1(e13, e13)) | (e16 = op1(e13, e12)) | (e16 = op1(e13, e11)) | (e16 = op1(e13, e10))) & ((e15 = op1(e16, e13)) | (e15 = op1(e15, e13)) | (e15 = op1(e14, e13)) | (e15 = op1(e13, e13)) | (e15 = op1(e12, e13)) | (e15 = op1(e11, e13)) | (e15 = op1(e10, e13))) & ((e15 = op1(e13, e16)) | (e15 = op1(e13, e15)) | (e15 = op1(e13, e14)) | (e15 = op1(e13, e13)) | (e15 = op1(e13, e12)) | (e15 = op1(e13, e11)) | (e15 = op1(e13, e10))) & ((e14 = op1(e16, e13)) | (e14 = op1(e15, e13)) | (e14 = op1(e14, e13)) | (e14 = op1(e13, e13)) | (e14 = op1(e12, e13)) | (e14 = op1(e11, e13)) | (e14 = op1(e10, e13))) & ((e14 = op1(e13, e16)) | (e14 = op1(e13, e15)) | (e14 = op1(e13, e14)) | (e14 = op1(e13, e13)) | (e14 = op1(e13, e12)) | (e14 = op1(e13, e11)) | (e14 = op1(e13, e10))) & ((e13 = op1(e16, e13)) | (e13 = op1(e15, e13)) | (e13 = op1(e14, e13)) | (e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e16)) | (e13 = op1(e13, e15)) | (e13 = op1(e13, e14)) | (e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e16, e13)) | (e12 = op1(e15, e13)) | (e12 = op1(e14, e13)) | (e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e16)) | (e12 = op1(e13, e15)) | (e12 = op1(e13, e14)) | (e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e16, e13)) | (e11 = op1(e15, e13)) | (e11 = op1(e14, e13)) | (e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e16)) | (e11 = op1(e13, e15)) | (e11 = op1(e13, e14)) | (e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e16, e13)) | (e10 = op1(e15, e13)) | (e10 = op1(e14, e13)) | (e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e16)) | (e10 = op1(e13, e15)) | (e10 = op1(e13, e14)) | (e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e16 = op1(e16, e12)) | (e16 = op1(e15, e12)) | (e16 = op1(e14, e12)) | (e16 = op1(e13, e12)) | (e16 = op1(e12, e12)) | (e16 = op1(e11, e12)) | (e16 = op1(e10, e12))) & ((e16 = op1(e12, e16)) | (e16 = op1(e12, e15)) | (e16 = op1(e12, e14)) | (e16 = op1(e12, e13)) | (e16 = op1(e12, e12)) | (e16 = op1(e12, e11)) | (e16 = op1(e12, e10))) & ((e15 = op1(e16, e12)) | (e15 = op1(e15, e12)) | (e15 = op1(e14, e12)) | (e15 = op1(e13, e12)) | (e15 = op1(e12, e12)) | (e15 = op1(e11, e12)) | (e15 = op1(e10, e12))) & ((e15 = op1(e12, e16)) | (e15 = op1(e12, e15)) | (e15 = op1(e12, e14)) | (e15 = op1(e12, e13)) | (e15 = op1(e12, e12)) | (e15 = op1(e12, e11)) | (e15 = op1(e12, e10))) & ((e14 = op1(e16, e12)) | (e14 = op1(e15, e12)) | (e14 = op1(e14, e12)) | (e14 = op1(e13, e12)) | (e14 = op1(e12, e12)) | (e14 = op1(e11, e12)) | (e14 = op1(e10, e12))) & ((e14 = op1(e12, e16)) | (e14 = op1(e12, e15)) | (e14 = op1(e12, e14)) | (e14 = op1(e12, e13)) | (e14 = op1(e12, e12)) | (e14 = op1(e12, e11)) | (e14 = op1(e12, e10))) & ((e13 = op1(e16, e12)) | (e13 = op1(e15, e12)) | (e13 = op1(e14, e12)) | (e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e16)) | (e13 = op1(e12, e15)) | (e13 = op1(e12, e14)) | (e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e16, e12)) | (e12 = op1(e15, e12)) | (e12 = op1(e14, e12)) | (e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e16)) | (e12 = op1(e12, e15)) | (e12 = op1(e12, e14)) | (e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e16, e12)) | (e11 = op1(e15, e12)) | (e11 = op1(e14, e12)) | (e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e16)) | (e11 = op1(e12, e15)) | (e11 = op1(e12, e14)) | (e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e16, e12)) | (e10 = op1(e15, e12)) | (e10 = op1(e14, e12)) | (e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e16)) | (e10 = op1(e12, e15)) | (e10 = op1(e12, e14)) | (e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e16 = op1(e16, e11)) | (e16 = op1(e15, e11)) | (e16 = op1(e14, e11)) | (e16 = op1(e13, e11)) | (e16 = op1(e12, e11)) | (e16 = op1(e11, e11)) | (e16 = op1(e10, e11))) & ((e16 = op1(e11, e16)) | (e16 = op1(e11, e15)) | (e16 = op1(e11, e14)) | (e16 = op1(e11, e13)) | (e16 = op1(e11, e12)) | (e16 = op1(e11, e11)) | (e16 = op1(e11, e10))) & ((e15 = op1(e16, e11)) | (e15 = op1(e15, e11)) | (e15 = op1(e14, e11)) | (e15 = op1(e13, e11)) | (e15 = op1(e12, e11)) | (e15 = op1(e11, e11)) | (e15 = op1(e10, e11))) & ((e15 = op1(e11, e16)) | (e15 = op1(e11, e15)) | (e15 = op1(e11, e14)) | (e15 = op1(e11, e13)) | (e15 = op1(e11, e12)) | (e15 = op1(e11, e11)) | (e15 = op1(e11, e10))) & ((e14 = op1(e16, e11)) | (e14 = op1(e15, e11)) | (e14 = op1(e14, e11)) | (e14 = op1(e13, e11)) | (e14 = op1(e12, e11)) | (e14 = op1(e11, e11)) | (e14 = op1(e10, e11))) & ((e14 = op1(e11, e16)) | (e14 = op1(e11, e15)) | (e14 = op1(e11, e14)) | (e14 = op1(e11, e13)) | (e14 = op1(e11, e12)) | (e14 = op1(e11, e11)) | (e14 = op1(e11, e10))) & ((e13 = op1(e16, e11)) | (e13 = op1(e15, e11)) | (e13 = op1(e14, e11)) | (e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e16)) | (e13 = op1(e11, e15)) | (e13 = op1(e11, e14)) | (e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e16, e11)) | (e12 = op1(e15, e11)) | (e12 = op1(e14, e11)) | (e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e16)) | (e12 = op1(e11, e15)) | (e12 = op1(e11, e14)) | (e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e16, e11)) | (e11 = op1(e15, e11)) | (e11 = op1(e14, e11)) | (e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e16)) | (e11 = op1(e11, e15)) | (e11 = op1(e11, e14)) | (e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e16, e11)) | (e10 = op1(e15, e11)) | (e10 = op1(e14, e11)) | (e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e16)) | (e10 = op1(e11, e15)) | (e10 = op1(e11, e14)) | (e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e16 = op1(e16, e10)) | (e16 = op1(e15, e10)) | (e16 = op1(e14, e10)) | (e16 = op1(e13, e10)) | (e16 = op1(e12, e10)) | (e16 = op1(e11, e10)) | (op1(e10, e10) = e16)) & ((e16 = op1(e10, e16)) | (e16 = op1(e10, e15)) | (e16 = op1(e10, e14)) | (e16 = op1(e10, e13)) | (e16 = op1(e10, e12)) | (e16 = op1(e10, e11)) | (op1(e10, e10) = e16)) & ((e15 = op1(e16, e10)) | (e15 = op1(e15, e10)) | (e15 = op1(e14, e10)) | (e15 = op1(e13, e10)) | (e15 = op1(e12, e10)) | (e15 = op1(e11, e10)) | (op1(e10, e10) = e15)) & ((e15 = op1(e10, e16)) | (e15 = op1(e10, e15)) | (e15 = op1(e10, e14)) | (e15 = op1(e10, e13)) | (e15 = op1(e10, e12)) | (e15 = op1(e10, e11)) | (op1(e10, e10) = e15)) & ((e14 = op1(e16, e10)) | (e14 = op1(e15, e10)) | (e14 = op1(e14, e10)) | (e14 = op1(e13, e10)) | (e14 = op1(e12, e10)) | (e14 = op1(e11, e10)) | (op1(e10, e10) = e14)) & ((e14 = op1(e10, e16)) | (e14 = op1(e10, e15)) | (e14 = op1(e10, e14)) | (e14 = op1(e10, e13)) | (e14 = op1(e10, e12)) | (e14 = op1(e10, e11)) | (op1(e10, e10) = e14)) & ((e13 = op1(e16, e10)) | (e13 = op1(e15, e10)) | (e13 = op1(e14, e10)) | (e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e16)) | (e13 = op1(e10, e15)) | (e13 = op1(e10, e14)) | (e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e16, e10)) | (e12 = op1(e15, e10)) | (e12 = op1(e14, e10)) | (e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e16)) | (e12 = op1(e10, e15)) | (e12 = op1(e10, e14)) | (e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e16, e10)) | (e11 = op1(e15, e10)) | (e11 = op1(e14, e10)) | (e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e16)) | (e11 = op1(e10, e15)) | (e11 = op1(e10, e14)) | (e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e16, e10)) | (e10 = op1(e15, e10)) | (e10 = op1(e14, e10)) | (e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e16)) | (e10 = op1(e10, e15)) | (e10 = op1(e10, e14)) | (e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax2)).
fof(f3117, plain, (spl42_337 | spl42_288 | spl42_239 | spl42_190 | spl42_141 | spl42_92 | spl42_43), inference(avatar_split_clause, [], [f161, f1775, f1978, f2181, f2384, f2587, f2790, f2993])).
fof(f161, plain, ((e10 = op1(e16, e10)) | (e10 = op1(e15, e10)) | (e10 = op1(e14, e10)) | (e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f3099, plain, (spl42_332 | spl42_283 | spl42_234 | spl42_185 | spl42_136 | spl42_87 | spl42_38), inference(avatar_split_clause, [], [f179, f1754, f1957, f2160, f2363, f2566, f2769, f2972])).
fof(f179, plain, ((e12 = op1(e16, e11)) | (e12 = op1(e15, e11)) | (e12 = op1(e14, e11)) | (e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f3094, plain, (spl42_293 | spl42_286 | spl42_279 | spl42_272 | spl42_265 | spl42_258 | spl42_251), inference(avatar_split_clause, [], [f184, f2636, f2665, f2694, f2723, f2752, f2781, f2810])).
fof(f184, plain, ((e15 = op1(e11, e16)) | (e15 = op1(e11, e15)) | (e15 = op1(e11, e14)) | (e15 = op1(e11, e13)) | (e15 = op1(e11, e12)) | (e15 = op1(e11, e11)) | (e15 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f3087, plain, (spl42_324 | spl42_275 | spl42_226 | spl42_177 | spl42_128 | spl42_79 | spl42_30), inference(avatar_split_clause, [], [f191, f1721, f1924, f2127, f2330, f2533, f2736, f2939])).
fof(f191, plain, ((e11 = op1(e16, e12)) | (e11 = op1(e15, e12)) | (e11 = op1(e14, e12)) | (e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f3082, plain, (spl42_243 | spl42_236 | spl42_229 | spl42_222 | spl42_215 | spl42_208 | spl42_201), inference(avatar_split_clause, [], [f196, f2429, f2458, f2487, f2516, f2545, f2574, f2603])).
fof(f196, plain, ((e14 = op1(e12, e16)) | (e14 = op1(e12, e15)) | (e14 = op1(e12, e14)) | (e14 = op1(e12, e13)) | (e14 = op1(e12, e12)) | (e14 = op1(e12, e11)) | (e14 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f3081, plain, (spl42_327 | spl42_278 | spl42_229 | spl42_180 | spl42_131 | spl42_82 | spl42_33), inference(avatar_split_clause, [], [f197, f1733, f1936, f2139, f2342, f2545, f2748, f2951])).
fof(f197, plain, ((e14 = op1(e16, e12)) | (e14 = op1(e15, e12)) | (e14 = op1(e14, e12)) | (e14 = op1(e13, e12)) | (e14 = op1(e12, e12)) | (e14 = op1(e11, e12)) | (e14 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f3078, plain, (spl42_245 | spl42_238 | spl42_231 | spl42_224 | spl42_217 | spl42_210 | spl42_203), inference(avatar_split_clause, [], [f200, f2437, f2466, f2495, f2524, f2553, f2582, f2611])).
fof(f200, plain, ((e16 = op1(e12, e16)) | (e16 = op1(e12, e15)) | (e16 = op1(e12, e14)) | (e16 = op1(e12, e13)) | (e16 = op1(e12, e12)) | (e16 = op1(e12, e11)) | (e16 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f3074, plain, (spl42_191 | spl42_184 | spl42_177 | spl42_170 | spl42_163 | spl42_156 | spl42_149), inference(avatar_split_clause, [], [f204, f2214, f2243, f2272, f2301, f2330, f2359, f2388])).
fof(f204, plain, ((e11 = op1(e13, e16)) | (e11 = op1(e13, e15)) | (e11 = op1(e13, e14)) | (e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f3073, plain, (spl42_317 | spl42_268 | spl42_219 | spl42_170 | spl42_121 | spl42_72 | spl42_23), inference(avatar_split_clause, [], [f205, f1692, f1895, f2098, f2301, f2504, f2707, f2910])).
fof(f205, plain, ((e11 = op1(e16, e13)) | (e11 = op1(e15, e13)) | (e11 = op1(e14, e13)) | (e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f3070, plain, (spl42_193 | spl42_186 | spl42_179 | spl42_172 | spl42_165 | spl42_158 | spl42_151), inference(avatar_split_clause, [], [f208, f2222, f2251, f2280, f2309, f2338, f2367, f2396])).
fof(f208, plain, ((e13 = op1(e13, e16)) | (e13 = op1(e13, e15)) | (e13 = op1(e13, e14)) | (e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f3067, plain, (spl42_320 | spl42_271 | spl42_222 | spl42_173 | spl42_124 | spl42_75 | spl42_26), inference(avatar_split_clause, [], [f211, f1704, f1907, f2110, f2313, f2516, f2719, f2922])).
fof(f211, plain, ((e14 = op1(e16, e13)) | (e14 = op1(e15, e13)) | (e14 = op1(e14, e13)) | (e14 = op1(e13, e13)) | (e14 = op1(e12, e13)) | (e14 = op1(e11, e13)) | (e14 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f3062, plain, (spl42_141 | spl42_134 | spl42_127 | spl42_120 | spl42_113 | spl42_106 | spl42_99), inference(avatar_split_clause, [], [f216, f2007, f2036, f2065, f2094, f2123, f2152, f2181])).
fof(f216, plain, ((e10 = op1(e14, e16)) | (e10 = op1(e14, e15)) | (e10 = op1(e14, e14)) | (e10 = op1(e14, e13)) | (e10 = op1(e14, e12)) | (e10 = op1(e14, e11)) | (e10 = op1(e14, e10))), inference(cnf_transformation, [], [f2])).
fof(f3056, plain, (spl42_144 | spl42_137 | spl42_130 | spl42_123 | spl42_116 | spl42_109 | spl42_102), inference(avatar_split_clause, [], [f222, f2019, f2048, f2077, f2106, f2135, f2164, f2193])).
fof(f222, plain, ((e13 = op1(e14, e16)) | (e13 = op1(e14, e15)) | (e13 = op1(e14, e14)) | (e13 = op1(e14, e13)) | (e13 = op1(e14, e12)) | (e13 = op1(e14, e11)) | (e13 = op1(e14, e10))), inference(cnf_transformation, [], [f2])).
fof(f3054, plain, (spl42_145 | spl42_138 | spl42_131 | spl42_124 | spl42_117 | spl42_110 | spl42_103), inference(avatar_split_clause, [], [f224, f2023, f2052, f2081, f2110, f2139, f2168, f2197])).
fof(f224, plain, ((e14 = op1(e14, e16)) | (e14 = op1(e14, e15)) | (e14 = op1(e14, e14)) | (e14 = op1(e14, e13)) | (e14 = op1(e14, e12)) | (e14 = op1(e14, e11)) | (e14 = op1(e14, e10))), inference(cnf_transformation, [], [f2])).
fof(f3048, plain, (spl42_92 | spl42_85 | spl42_78 | spl42_71 | spl42_64 | spl42_57 | spl42_50), inference(avatar_split_clause, [], [f230, f1804, f1833, f1862, f1891, f1920, f1949, f1978])).
fof(f230, plain, ((e10 = op1(e15, e16)) | (e10 = op1(e15, e15)) | (e10 = op1(e15, e14)) | (e10 = op1(e15, e13)) | (e10 = op1(e15, e12)) | (e10 = op1(e15, e11)) | (e10 = op1(e15, e10))), inference(cnf_transformation, [], [f2])).
fof(f3044, plain, (spl42_94 | spl42_87 | spl42_80 | spl42_73 | spl42_66 | spl42_59 | spl42_52), inference(avatar_split_clause, [], [f234, f1812, f1841, f1870, f1899, f1928, f1957, f1986])).
fof(f234, plain, ((e12 = op1(e15, e16)) | (e12 = op1(e15, e15)) | (e12 = op1(e15, e14)) | (e12 = op1(e15, e13)) | (e12 = op1(e15, e12)) | (e12 = op1(e15, e11)) | (e12 = op1(e15, e10))), inference(cnf_transformation, [], [f2])).
fof(f3035, plain, (spl42_308 | spl42_259 | spl42_210 | spl42_161 | spl42_112 | spl42_63 | spl42_14), inference(avatar_split_clause, [], [f243, f1654, f1857, f2060, f2263, f2466, f2669, f2872])).
fof(f243, plain, ((e16 = op1(e16, e15)) | (e16 = op1(e15, e15)) | (e16 = op1(e14, e15)) | (e16 = op1(e13, e15)) | (e16 = op1(e12, e15)) | (e16 = op1(e11, e15)) | (e16 = op1(e10, e15))), inference(cnf_transformation, [], [f2])).
fof(f3023, plain, (spl42_300 | spl42_251 | spl42_202 | spl42_153 | spl42_104 | spl42_55 | spl42_6), inference(avatar_split_clause, [], [f255, f1621, f1824, f2027, f2230, f2433, f2636, f2839])).
fof(f255, plain, ((e15 = op1(e16, e16)) | (e15 = op1(e15, e16)) | (e15 = op1(e14, e16)) | (e15 = op1(e13, e16)) | (e15 = op1(e12, e16)) | (e15 = op1(e11, e16)) | (e15 = op1(e10, e16))), inference(cnf_transformation, [], [f2])).
fof(f2962, plain, (spl42_323 | spl42_324 | spl42_325 | spl42_326 | spl42_327 | spl42_328 | spl42_329), inference(avatar_split_clause, [], [f113, f2959, f2955, f2951, f2947, f2943, f2939, f2935])).
fof(f113, plain, ((e16 = op1(e10, e12)) | (e15 = op1(e10, e12)) | (e14 = op1(e10, e12)) | (e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e16 = op1(e16, e16)) | (e15 = op1(e16, e16)) | (e14 = op1(e16, e16)) | (e13 = op1(e16, e16)) | (e12 = op1(e16, e16)) | (e11 = op1(e16, e16)) | (e10 = op1(e16, e16))) & ((e16 = op1(e16, e15)) | (e15 = op1(e16, e15)) | (e14 = op1(e16, e15)) | (e13 = op1(e16, e15)) | (e12 = op1(e16, e15)) | (e11 = op1(e16, e15)) | (e10 = op1(e16, e15))) & ((e16 = op1(e16, e14)) | (e15 = op1(e16, e14)) | (e14 = op1(e16, e14)) | (e13 = op1(e16, e14)) | (e12 = op1(e16, e14)) | (e11 = op1(e16, e14)) | (e10 = op1(e16, e14))) & ((e16 = op1(e16, e13)) | (e15 = op1(e16, e13)) | (e14 = op1(e16, e13)) | (e13 = op1(e16, e13)) | (e12 = op1(e16, e13)) | (e11 = op1(e16, e13)) | (e10 = op1(e16, e13))) & ((e16 = op1(e16, e12)) | (e15 = op1(e16, e12)) | (e14 = op1(e16, e12)) | (e13 = op1(e16, e12)) | (e12 = op1(e16, e12)) | (e11 = op1(e16, e12)) | (e10 = op1(e16, e12))) & ((e16 = op1(e16, e11)) | (e15 = op1(e16, e11)) | (e14 = op1(e16, e11)) | (e13 = op1(e16, e11)) | (e12 = op1(e16, e11)) | (e11 = op1(e16, e11)) | (e10 = op1(e16, e11))) & ((e16 = op1(e16, e10)) | (e15 = op1(e16, e10)) | (e14 = op1(e16, e10)) | (e13 = op1(e16, e10)) | (e12 = op1(e16, e10)) | (e11 = op1(e16, e10)) | (e10 = op1(e16, e10))) & ((e16 = op1(e15, e16)) | (e15 = op1(e15, e16)) | (e14 = op1(e15, e16)) | (e13 = op1(e15, e16)) | (e12 = op1(e15, e16)) | (e11 = op1(e15, e16)) | (e10 = op1(e15, e16))) & ((e16 = op1(e15, e15)) | (e15 = op1(e15, e15)) | (e14 = op1(e15, e15)) | (e13 = op1(e15, e15)) | (e12 = op1(e15, e15)) | (e11 = op1(e15, e15)) | (e10 = op1(e15, e15))) & ((e16 = op1(e15, e14)) | (e15 = op1(e15, e14)) | (e14 = op1(e15, e14)) | (e13 = op1(e15, e14)) | (e12 = op1(e15, e14)) | (e11 = op1(e15, e14)) | (e10 = op1(e15, e14))) & ((e16 = op1(e15, e13)) | (e15 = op1(e15, e13)) | (e14 = op1(e15, e13)) | (e13 = op1(e15, e13)) | (e12 = op1(e15, e13)) | (e11 = op1(e15, e13)) | (e10 = op1(e15, e13))) & ((e16 = op1(e15, e12)) | (e15 = op1(e15, e12)) | (e14 = op1(e15, e12)) | (e13 = op1(e15, e12)) | (e12 = op1(e15, e12)) | (e11 = op1(e15, e12)) | (e10 = op1(e15, e12))) & ((e16 = op1(e15, e11)) | (e15 = op1(e15, e11)) | (e14 = op1(e15, e11)) | (e13 = op1(e15, e11)) | (e12 = op1(e15, e11)) | (e11 = op1(e15, e11)) | (e10 = op1(e15, e11))) & ((e16 = op1(e15, e10)) | (e15 = op1(e15, e10)) | (e14 = op1(e15, e10)) | (e13 = op1(e15, e10)) | (e12 = op1(e15, e10)) | (e11 = op1(e15, e10)) | (e10 = op1(e15, e10))) & ((e16 = op1(e14, e16)) | (e15 = op1(e14, e16)) | (e14 = op1(e14, e16)) | (e13 = op1(e14, e16)) | (e12 = op1(e14, e16)) | (e11 = op1(e14, e16)) | (e10 = op1(e14, e16))) & ((e16 = op1(e14, e15)) | (e15 = op1(e14, e15)) | (e14 = op1(e14, e15)) | (e13 = op1(e14, e15)) | (e12 = op1(e14, e15)) | (e11 = op1(e14, e15)) | (e10 = op1(e14, e15))) & ((e16 = op1(e14, e14)) | (e15 = op1(e14, e14)) | (e14 = op1(e14, e14)) | (e13 = op1(e14, e14)) | (e12 = op1(e14, e14)) | (e11 = op1(e14, e14)) | (e10 = op1(e14, e14))) & ((e16 = op1(e14, e13)) | (e15 = op1(e14, e13)) | (e14 = op1(e14, e13)) | (e13 = op1(e14, e13)) | (e12 = op1(e14, e13)) | (e11 = op1(e14, e13)) | (e10 = op1(e14, e13))) & ((e16 = op1(e14, e12)) | (e15 = op1(e14, e12)) | (e14 = op1(e14, e12)) | (e13 = op1(e14, e12)) | (e12 = op1(e14, e12)) | (e11 = op1(e14, e12)) | (e10 = op1(e14, e12))) & ((e16 = op1(e14, e11)) | (e15 = op1(e14, e11)) | (e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))) & ((e16 = op1(e14, e10)) | (e15 = op1(e14, e10)) | (e14 = op1(e14, e10)) | (e13 = op1(e14, e10)) | (e12 = op1(e14, e10)) | (e11 = op1(e14, e10)) | (e10 = op1(e14, e10))) & ((e16 = op1(e13, e16)) | (e15 = op1(e13, e16)) | (e14 = op1(e13, e16)) | (e13 = op1(e13, e16)) | (e12 = op1(e13, e16)) | (e11 = op1(e13, e16)) | (e10 = op1(e13, e16))) & ((e16 = op1(e13, e15)) | (e15 = op1(e13, e15)) | (e14 = op1(e13, e15)) | (e13 = op1(e13, e15)) | (e12 = op1(e13, e15)) | (e11 = op1(e13, e15)) | (e10 = op1(e13, e15))) & ((e16 = op1(e13, e14)) | (e15 = op1(e13, e14)) | (e14 = op1(e13, e14)) | (e13 = op1(e13, e14)) | (e12 = op1(e13, e14)) | (e11 = op1(e13, e14)) | (e10 = op1(e13, e14))) & ((e16 = op1(e13, e13)) | (e15 = op1(e13, e13)) | (e14 = op1(e13, e13)) | (e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e16 = op1(e13, e12)) | (e15 = op1(e13, e12)) | (e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e16 = op1(e13, e11)) | (e15 = op1(e13, e11)) | (e14 = op1(e13, e11)) | (e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e16 = op1(e13, e10)) | (e15 = op1(e13, e10)) | (e14 = op1(e13, e10)) | (e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e16 = op1(e12, e16)) | (e15 = op1(e12, e16)) | (e14 = op1(e12, e16)) | (e13 = op1(e12, e16)) | (e12 = op1(e12, e16)) | (e11 = op1(e12, e16)) | (e10 = op1(e12, e16))) & ((e16 = op1(e12, e15)) | (e15 = op1(e12, e15)) | (e14 = op1(e12, e15)) | (e13 = op1(e12, e15)) | (e12 = op1(e12, e15)) | (e11 = op1(e12, e15)) | (e10 = op1(e12, e15))) & ((e16 = op1(e12, e14)) | (e15 = op1(e12, e14)) | (e14 = op1(e12, e14)) | (e13 = op1(e12, e14)) | (e12 = op1(e12, e14)) | (e11 = op1(e12, e14)) | (e10 = op1(e12, e14))) & ((e16 = op1(e12, e13)) | (e15 = op1(e12, e13)) | (e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e16 = op1(e12, e12)) | (e15 = op1(e12, e12)) | (e14 = op1(e12, e12)) | (e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e16 = op1(e12, e11)) | (e15 = op1(e12, e11)) | (e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e16 = op1(e12, e10)) | (e15 = op1(e12, e10)) | (e14 = op1(e12, e10)) | (e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e16 = op1(e11, e16)) | (e15 = op1(e11, e16)) | (e14 = op1(e11, e16)) | (e13 = op1(e11, e16)) | (e12 = op1(e11, e16)) | (e11 = op1(e11, e16)) | (e10 = op1(e11, e16))) & ((e16 = op1(e11, e15)) | (e15 = op1(e11, e15)) | (e14 = op1(e11, e15)) | (e13 = op1(e11, e15)) | (e12 = op1(e11, e15)) | (e11 = op1(e11, e15)) | (e10 = op1(e11, e15))) & ((e16 = op1(e11, e14)) | (e15 = op1(e11, e14)) | (e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))) & ((e16 = op1(e11, e13)) | (e15 = op1(e11, e13)) | (e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e16 = op1(e11, e12)) | (e15 = op1(e11, e12)) | (e14 = op1(e11, e12)) | (e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e16 = op1(e11, e11)) | (e15 = op1(e11, e11)) | (e14 = op1(e11, e11)) | (e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e16 = op1(e11, e10)) | (e15 = op1(e11, e10)) | (e14 = op1(e11, e10)) | (e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e16 = op1(e10, e16)) | (e15 = op1(e10, e16)) | (e14 = op1(e10, e16)) | (e13 = op1(e10, e16)) | (e12 = op1(e10, e16)) | (e11 = op1(e10, e16)) | (e10 = op1(e10, e16))) & ((e16 = op1(e10, e15)) | (e15 = op1(e10, e15)) | (e14 = op1(e10, e15)) | (e13 = op1(e10, e15)) | (e12 = op1(e10, e15)) | (e11 = op1(e10, e15)) | (e10 = op1(e10, e15))) & ((e16 = op1(e10, e14)) | (e15 = op1(e10, e14)) | (e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))) & ((e16 = op1(e10, e13)) | (e15 = op1(e10, e13)) | (e14 = op1(e10, e13)) | (e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e16 = op1(e10, e12)) | (e15 = op1(e10, e12)) | (e14 = op1(e10, e12)) | (e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e16 = op1(e10, e11)) | (e15 = op1(e10, e11)) | (e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e16) | (op1(e10, e10) = e15) | (op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG195+1.p', ax1)).
fof(f2701, plain, (spl42_260 | spl42_261 | spl42_262 | spl42_263 | spl42_264 | spl42_265 | spl42_266), inference(avatar_split_clause, [], [f122, f2698, f2694, f2690, f2686, f2682, f2678, f2674])).
fof(f122, plain, ((e16 = op1(e11, e14)) | (e15 = op1(e11, e14)) | (e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))), inference(cnf_transformation, [], [f1])).
fof(f2585, plain, (spl42_232 | spl42_233 | spl42_234 | spl42_235 | spl42_236 | spl42_237 | spl42_238), inference(avatar_split_clause, [], [f126, f2582, f2578, f2574, f2570, f2566, f2562, f2558])).
fof(f126, plain, ((e16 = op1(e12, e11)) | (e15 = op1(e12, e11)) | (e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f2527, plain, (spl42_218 | spl42_219 | spl42_220 | spl42_221 | spl42_222 | spl42_223 | spl42_224), inference(avatar_split_clause, [], [f128, f2524, f2520, f2516, f2512, f2508, f2504, f2500])).
fof(f128, plain, ((e16 = op1(e12, e13)) | (e15 = op1(e12, e13)) | (e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))), inference(cnf_transformation, [], [f1])).
fof(f2353, plain, (spl42_176 | spl42_177 | spl42_178 | spl42_179 | spl42_180 | spl42_181 | spl42_182), inference(avatar_split_clause, [], [f134, f2350, f2346, f2342, f2338, f2334, f2330, f2326])).
fof(f134, plain, ((e16 = op1(e13, e12)) | (e15 = op1(e13, e12)) | (e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))), inference(cnf_transformation, [], [f1])).
fof(f2266, plain, (spl42_155 | spl42_156 | spl42_157 | spl42_158 | spl42_159 | spl42_160 | spl42_161), inference(avatar_split_clause, [], [f137, f2263, f2259, f2255, f2251, f2247, f2243, f2239])).
fof(f137, plain, ((e16 = op1(e13, e15)) | (e15 = op1(e13, e15)) | (e14 = op1(e13, e15)) | (e13 = op1(e13, e15)) | (e12 = op1(e13, e15)) | (e11 = op1(e13, e15)) | (e10 = op1(e13, e15))), inference(cnf_transformation, [], [f1])).
fof(f2063, plain, (spl42_106 | spl42_107 | spl42_108 | spl42_109 | spl42_110 | spl42_111 | spl42_112), inference(avatar_split_clause, [], [f144, f2060, f2056, f2052, f2048, f2044, f2040, f2036])).
fof(f144, plain, ((e16 = op1(e14, e15)) | (e15 = op1(e14, e15)) | (e14 = op1(e14, e15)) | (e13 = op1(e14, e15)) | (e12 = op1(e14, e15)) | (e11 = op1(e14, e15)) | (e10 = op1(e14, e15))), inference(cnf_transformation, [], [f1])).
fof(f2034, plain, (spl42_99 | spl42_100 | spl42_101 | spl42_102 | spl42_103 | spl42_104 | spl42_105), inference(avatar_split_clause, [], [f145, f2031, f2027, f2023, f2019, f2015, f2011, f2007])).
fof(f145, plain, ((e16 = op1(e14, e16)) | (e15 = op1(e14, e16)) | (e14 = op1(e14, e16)) | (e13 = op1(e14, e16)) | (e12 = op1(e14, e16)) | (e11 = op1(e14, e16)) | (e10 = op1(e14, e16))), inference(cnf_transformation, [], [f1])).
fof(f1947, plain, (spl42_78 | spl42_79 | spl42_80 | spl42_81 | spl42_82 | spl42_83 | spl42_84), inference(avatar_split_clause, [], [f148, f1944, f1940, f1936, f1932, f1928, f1924, f1920])).
fof(f148, plain, ((e16 = op1(e15, e12)) | (e15 = op1(e15, e12)) | (e14 = op1(e15, e12)) | (e13 = op1(e15, e12)) | (e12 = op1(e15, e12)) | (e11 = op1(e15, e12)) | (e10 = op1(e15, e12))), inference(cnf_transformation, [], [f1])).
fof(f1918, plain, (spl42_71 | spl42_72 | spl42_73 | spl42_74 | spl42_75 | spl42_76 | spl42_77), inference(avatar_split_clause, [], [f149, f1915, f1911, f1907, f1903, f1899, f1895, f1891])).
fof(f149, plain, ((e16 = op1(e15, e13)) | (e15 = op1(e15, e13)) | (e14 = op1(e15, e13)) | (e13 = op1(e15, e13)) | (e12 = op1(e15, e13)) | (e11 = op1(e15, e13)) | (e10 = op1(e15, e13))), inference(cnf_transformation, [], [f1])).
fof(f1889, plain, (spl42_64 | spl42_65 | spl42_66 | spl42_67 | spl42_68 | spl42_69 | spl42_70), inference(avatar_split_clause, [], [f150, f1886, f1882, f1878, f1874, f1870, f1866, f1862])).
fof(f150, plain, ((e16 = op1(e15, e14)) | (e15 = op1(e15, e14)) | (e14 = op1(e15, e14)) | (e13 = op1(e15, e14)) | (e12 = op1(e15, e14)) | (e11 = op1(e15, e14)) | (e10 = op1(e15, e14))), inference(cnf_transformation, [], [f1])).
fof(f1715, plain, (spl42_22 | spl42_23 | spl42_24 | spl42_25 | spl42_26 | spl42_27 | spl42_28), inference(avatar_split_clause, [], [f156, f1712, f1708, f1704, f1700, f1696, f1692, f1688])).
fof(f156, plain, ((e16 = op1(e16, e13)) | (e15 = op1(e16, e13)) | (e14 = op1(e16, e13)) | (e13 = op1(e16, e13)) | (e12 = op1(e16, e13)) | (e11 = op1(e16, e13)) | (e10 = op1(e16, e13))), inference(cnf_transformation, [], [f1])).
fof(f1686, plain, (spl42_15 | spl42_16 | spl42_17 | spl42_18 | spl42_19 | spl42_20 | spl42_21), inference(avatar_split_clause, [], [f157, f1683, f1679, f1675, f1671, f1667, f1663, f1659])).
fof(f157, plain, ((e16 = op1(e16, e14)) | (e15 = op1(e16, e14)) | (e14 = op1(e16, e14)) | (e13 = op1(e16, e14)) | (e12 = op1(e16, e14)) | (e11 = op1(e16, e14)) | (e10 = op1(e16, e14))), inference(cnf_transformation, [], [f1])).