fof(f9776, plain, $false, inference(avatar_sat_refutation, [], [f707, f728, f749, f770, f791, f812, f854, f875, f896, f917, f938, f959, f980, f1001, f1022, f1043, f1064, f1085, f1106, f1127, f1148, f1169, f1190, f1211, f1232, f1235, f1236, f1237, f1238, f1239, f1240, f1241, f1242, f1243, f1244, f1247, f1248, f1249, f1250, f1251, f1252, f1253, f1254, f1255, f1256, f1259, f1260, f1261, f1262, f1263, f1264, f1265, f1266, f1267, f1268, f1271, f1272, f1273, f1274, f1275, f1276, f1277, f1278, f1279, f1280, f1287, f1288, f1289, f1294, f1295, f1296, f1301, f1302, f1303, f1309, f1310, f1315, f1316, f1317, f1322, f1324, f1329, f1330, f1331, f1336, f1337, f1338, f1343, f1344, f1345, f1350, f1351, f1352, f1357, f1358, f1359, f1364, f1366, f1371, f1372, f1373, f1378, f1379, f1380, f1385, f1386, f1387, f1392, f1393, f1394, f1399, f1400, f1401, f1406, f1408, f1413, f1414, f1415, f1420, f1421, f1422, f1427, f1428, f1429, f1434, f1435, f1436, f1441, f1442, f1443, f1448, f1450, f1451, f1453, f1501, f1519, f1547, f1568, f1573, f1584, f1585, f1587, f1589, f1600, f1605, f1606, f1607, f1622, f1649, f1656, f1661, f1666, f1671, f1672, f1673, f1674, f1676, f1682, f1688, f1693, f1695, f1704, f1711, f1717, f1722, f1727, f1728, f1730, f1733, f1738, f1743, f1744, f1749, f1751, f1753, f1760, f1763, f1765, f1766, f1768, f1769, f1770, f1771, f1772, f1773, f1783, f1789, f1790, f1791, f1792, f1800, f1804, f1806, f1807, f1808, f1809, f1810, f1811, f1812, f1813, f1849, f1853, f1866, f1868, f1869, f1871, f1880, f1892, f1898, f1904, f1907, f1910, f1918, f1920, f1922, f1935, f1940, f1951, f1971, f1974, f1985, f1986, f1989, f1995, f1997, f1999, f2009, f2017, f2018, f2023, f2024, f2040, f2047, f2052, f2071, f2074, f2077, f2085, f2088, f2091, f2094, f2103, f2106, f2109, f2113, f2125, f2128, f2141, f2153, f2164, f2168, f2178, f2193, f2258, f2274, f2281, f2284, f2287, f2301, f2315, f2323, f2324, f2335, f2348, f2379, f2382, f2388, f2396, f2399, f2414, f2448, f2460, f2485, f2489, f2492, f2494, f2506, f2517, f2523, f2546, f2547, f2555, f2563, f2564, f2583, f2592, f2607, f2616, f2622, f2627, f2673, f2677, f2683, f2695, f2713, f2728, f2734, f2752, f2755, f2757, f2758, f2771, f2796, f2804, f2815, f2829, f2876, f2894, f2935, f2938, f2943, f2946, f2977, f2989, f3058, f3062, f3065, f3072, f3075, f3108, f3142, f3145, f3148, f3151, f3154, f3177, f3206, f3232, f3236, f3242, f3258, f3320, f3326, f3329, f3332, f3336, f3374, f3379, f3402, f3411, f3431, f3487, f3537, f3570, f3589, f3605, f3613, f3618, f3621, f3625, f3628, f3635, f3638, f3641, f3653, f3668, f3675, f3688, f3689, f3696, f3710, f3729, f3732, f3735, f3781, f3791, f3826, f3860, f3862, f3865, f3868, f3872, f3888, f3890, f3896, f3897, f3908, f3913, f3914, f3917, f3943, f3979, f3982, f3993, f4000, f4041, f4061, f4067, f4077, f4081, f4085, f4088, f4090, f4108, f4117, f4123, f4129, f4130, f4138, f4144, f4195, f4210, f4211, f4218, f4230, f4231, f4234, f4239, f4269, f4296, f4307, f4313, f4315, f4329, f4331, f4383, f4385, f4402, f4410, f4421, f4443, f4450, f4452, f4454, f4458, f4460, f4464, f4471, f4524, f4545, f4548, f4567, f4578, f4580, f4582, f4613, f4614, f4621, f4624, f4643, f4644, f4646, f4651, f4665, f4682, f4685, f4687, f4715, f4720, f4722, f4723, f4732, f4780, f4797, f4801, f4828, f4852, f4899, f4901, f4904, f4911, f4921, f4939, f4943, f4948, f4980, f4984, f4987, f5015, f5025, f5075, f5099, f5125, f5133, f5136, f5156, f5207, f5226, f5235, f5238, f5287, f5309, f5333, f5371, f5405, f5433, f5447, f5467, f5493, f5494, f5496, f5497, f5498, f5499, f5500, f5501, f5502, f5513, f5514, f5523, f5533, f5549, f5550, f5569, f5593, f5599, f5614, f5615, f5637, f5641, f5643, f5650, f5669, f5672, f5685, f5694, f5697, f5700, f5717, f5725, f5730, f5767, f5778, f5779, f5796, f5798, f5804, f5806, f5818, f5829, f5842, f5846, f5850, f5857, f5859, f5867, f5886, f5909, f5914, f5919, f5925, f5936, f5950, f5966, f5977, f5984, f5994, f6000, f6019, f6021, f6023, f6032, f6047, f6090, f6091, f6117, f6139, f6140, f6141, f6159, f6168, f6223, f6233, f6248, f6250, f6298, f6331, f6345, f6346, f6350, f6356, f6388, f6411, f6424, f6425, f6434, f6487, f6488, f6490, f6503, f6514, f6535, f6549, f6583, f6584, f6608, f6614, f6630, f6639, f6654, f6655, f6656, f6680, f6688, f6727, f6741, f6750, f6752, f6760, f6784, f6810, f6812, f6840, f6856, f6868, f6877, f6885, f6886, f6907, f6951, f6952, f6958, f6959, f6973, f6974, f6987, f7048, f7054, f7063, f7068, f7070, f7079, f7106, f7119, f7120, f7129, f7138, f7139, f7145, f7168, f7172, f7200, f7214, f7225, f7231, f7241, f7266, f7270, f7272, f7283, f7311, f7323, f7327, f7331, f7333, f7337, f7346, f7347, f7353, f7366, f7381, f7382, f7383, f7384, f7398, f7406, f7435, f7439, f7441, f7443, f7448, f7450, f7452, f7458, f7464, f7511, f7530, f7531, f7537, f7549, f7552, f7561, f7568, f7586, f7587, f7589, f7607, f7608, f7616, f7629, f7652, f7695, f7698, f7709, f7727, f7752, f7761, f7773, f7777, f7788, f7796, f7803, f7806, f7810, f7817, f7822, f7837, f7841, f7845, f7854, f7856, f7860, f7879, f7881, f7883, f7887, f7893, f7921, f7930, f7956, f7972, f7980, f7999, f8000, f8095, f8111, f8130, f8131, f8132, f8133, f8135, f8136, f8137, f8138, f8139, f8154, f8155, f8156, f8164, f8174, f8175, f8182, f8183, f8192, f8215, f8230, f8231, f8240, f8256, f8257, f8277, f8283, f8292, f8297, f8307, f8311, f8319, f8325, f8328, f8344, f8355, f8368, f8374, f8398, f8401, f8407, f8418, f8425, f8426, f8440, f8449, f8450, f8458, f8467, f8477, f8478, f8479, f8489, f8508, f8516, f8519, f8526, f8528, f8530, f8535, f8544, f8583, f8584, f8590, f8606, f8622, f8623, f8627, f8629, f8635, f8650, f8651, f8657, f8682, f8700, f8730, f8732, f8750, f8755, f8865, f8866, f8867, f8868, f8877, f8881, f8906, f8907, f8942, f8962, f8967, f8973, f8986, f9012, f9017, f9021, f9043, f9044, f9045, f9046, f9047, f9048, f9049, f9050, f9052, f9069, f9070, f9071, f9073, f9082, f9083, f9089, f9090, f9099, f9108, f9117, f9135, f9136, f9137, f9139, f9140, f9149, f9160, f9175, f9185, f9186, f9202, f9207, f9215, f9224, f9226, f9237, f9244, f9264, f9269, f9281, f9283, f9287, f9301, f9319, f9320, f9321, f9322, f9323, f9324, f9326, f9327, f9328, f9334, f9335, f9346, f9359, f9360, f9368, f9376, f9377, f9386, f9394, f9406, f9413, f9417, f9427, f9436, f9442, f9443, f9456, f9457, f9474, f9479, f9487, f9495, f9499, f9533, f9542, f9560, f9561, f9562, f9563, f9564, f9565, f9566, f9567, f9568, f9592, f9617, f9618, f9627, f9634, f9637, f9648, f9649, f9664, f9665, f9667, f9668, f9678, f9679, f9686, f9693, f9701, f9705, f9714, f9716, f9726, f9733, f9739, f9751, f9752, f9761, f9766])).
fof(f9766, plain, (~ spl24_16 | ~ spl24_19), inference(avatar_contradiction_clause, [], [f9765])).
fof(f9765, plain, ($false | (~ spl24_16 | ~ spl24_19)), inference(subsumption_resolution, [], [f9764, f484])).
fof(f484, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (e3 = e4) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax5)).
fof(f9764, plain, ((e0 = e3) | (~ spl24_16 | ~ spl24_19)), inference(backward_demodulation, [], [f765, f753])).
fof(f753, plain, ((e0 = op(e4, e1)) | ~ spl24_16), inference(avatar_component_clause, [], [f751])).
fof(f751, plain, (spl24_16 <=> (e0 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_16])])).
fof(f765, plain, ((e3 = op(e4, e1)) | ~ spl24_19), inference(avatar_component_clause, [], [f763])).
fof(f763, plain, (spl24_19 <=> (e3 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_19])])).
fof(f9761, plain, (~ spl24_10 | ~ spl24_25), inference(avatar_split_clause, [], [f9760, f788, f725])).
fof(f725, plain, (spl24_10 <=> (e4 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_10])])).
fof(f788, plain, (spl24_25 <=> (e4 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_25])])).
fof(f9760, plain, (~ (e4 = op(e4, e3)) | ~ spl24_25), inference(backward_demodulation, [], [f474, f790])).
fof(f790, plain, ((e4 = op(e4, e0)) | ~ spl24_25), inference(avatar_component_clause, [], [f788])).
fof(f474, plain, ~ (op(e4, e0) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (op(e4, e3) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e4)) & ~ (op(e4, e2) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e4)) & ~ (op(e4, e1) = op(e4, e3)) & ~ (op(e4, e1) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e4)) & ~ (op(e4, e0) = op(e4, e3)) & ~ (op(e4, e0) = op(e4, e2)) & ~ (op(e4, e0) = op(e4, e1)) & ~ (op(e3, e3) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e4)) & ~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e4)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e4)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e3) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e4)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e4)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e4)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e3) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e4)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e4)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e4)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e3) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e4)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e4)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e4)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e3, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e4, e4)) & ~ (op(e2, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e4, e4)) & ~ (op(e1, e4) = op(e3, e4)) & ~ (op(e1, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e4, e4)) & ~ (op(e0, e4) = op(e3, e4)) & ~ (op(e0, e4) = op(e2, e4)) & ~ (op(e0, e4) = op(e1, e4)) & ~ (op(e3, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e4, e3)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e4, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e4, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e3, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e4, e2)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e4, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e4, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e3, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e4, e1)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e4, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e4, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e3, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e4, e0)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e4, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e4, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax4)).
fof(f9752, plain, (~ spl24_30 | ~ spl24_35), inference(avatar_split_clause, [], [f9747, f830, f809])).
fof(f809, plain, (spl24_30 <=> (e4 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_30])])).
fof(f830, plain, (spl24_35 <=> (e4 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_35])])).
fof(f9747, plain, (~ (e4 = op(e3, e4)) | ~ spl24_35), inference(backward_demodulation, [], [f471, f832])).
fof(f832, plain, ((e4 = op(e3, e3)) | ~ spl24_35), inference(avatar_component_clause, [], [f830])).
fof(f471, plain, ~ (op(e3, e3) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f9751, plain, (~ spl24_10 | ~ spl24_35), inference(avatar_split_clause, [], [f9746, f830, f725])).
fof(f9746, plain, (~ (e4 = op(e4, e3)) | ~ spl24_35), inference(backward_demodulation, [], [f421, f832])).
fof(f421, plain, ~ (op(e3, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f9739, plain, (~ spl24_38 | ~ spl24_43), inference(avatar_split_clause, [], [f9734, f864, f843])).
fof(f843, plain, (spl24_38 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_38])])).
fof(f864, plain, (spl24_43 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_43])])).
fof(f9734, plain, (~ (e2 = op(e3, e2)) | ~ spl24_43), inference(backward_demodulation, [], [f466, f866])).
fof(f866, plain, ((e2 = op(e3, e1)) | ~ spl24_43), inference(avatar_component_clause, [], [f864])).
fof(f466, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f9733, plain, (~ spl24_34 | ~ spl24_49), inference(avatar_split_clause, [], [f9730, f889, f826])).
fof(f826, plain, (spl24_34 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_34])])).
fof(f889, plain, (spl24_49 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_49])])).
fof(f9730, plain, (~ (e3 = op(e3, e3)) | ~ spl24_49), inference(backward_demodulation, [], [f464, f891])).
fof(f891, plain, ((e3 = op(e3, e0)) | ~ spl24_49), inference(avatar_component_clause, [], [f889])).
fof(f464, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f9726, plain, (~ spl24_52 | ~ spl24_54), inference(avatar_contradiction_clause, [], [f9725])).
fof(f9725, plain, ($false | (~ spl24_52 | ~ spl24_54)), inference(subsumption_resolution, [], [f9724, f487])).
fof(f487, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f5])).
fof(f9724, plain, ((e1 = e3) | (~ spl24_52 | ~ spl24_54)), inference(backward_demodulation, [], [f912, f904])).
fof(f904, plain, ((e1 = op(e2, e4)) | ~ spl24_52), inference(avatar_component_clause, [], [f902])).
fof(f902, plain, (spl24_52 <=> (e1 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_52])])).
fof(f912, plain, ((e3 = op(e2, e4)) | ~ spl24_54), inference(avatar_component_clause, [], [f910])).
fof(f910, plain, (spl24_54 <=> (e3 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_54])])).
fof(f9716, plain, (spl24_14 | ~ spl24_65 | ~ spl24_208), inference(avatar_split_clause, [], [f9713, f1740, f956, f742])).
fof(f742, plain, (spl24_14 <=> (e3 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_14])])).
fof(f956, plain, (spl24_65 <=> (e4 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_65])])).
fof(f1740, plain, (spl24_208 <=> (e3 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_208])])).
fof(f9713, plain, ((e3 = op(e4, e2)) | (~ spl24_65 | ~ spl24_208)), inference(backward_demodulation, [], [f1741, f958])).
fof(f958, plain, ((e4 = op(e2, e2)) | ~ spl24_65), inference(avatar_component_clause, [], [f956])).
fof(f1741, plain, ((e3 = op(op(e2, e2), e2)) | ~ spl24_208), inference(avatar_component_clause, [], [f1740])).
fof(f9714, plain, (~ spl24_11 | ~ spl24_65 | spl24_163), inference(avatar_split_clause, [], [f9709, f1490, f956, f730])).
fof(f730, plain, (spl24_11 <=> (e0 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_11])])).
fof(f1490, plain, (spl24_163 <=> (e0 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_163])])).
fof(f9709, plain, (~ (e0 = op(e4, e2)) | (~ spl24_65 | spl24_163)), inference(backward_demodulation, [], [f1492, f958])).
fof(f1492, plain, (~ (e0 = op(op(e2, e2), e2)) | spl24_163), inference(avatar_component_clause, [], [f1490])).
fof(f9705, plain, (~ spl24_65 | ~ spl24_70), inference(avatar_split_clause, [], [f9703, f977, f956])).
fof(f977, plain, (spl24_70 <=> (e4 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_70])])).
fof(f9703, plain, (~ (e4 = op(e2, e2)) | ~ spl24_70), inference(backward_demodulation, [], [f456, f979])).
fof(f979, plain, ((e4 = op(e2, e1)) | ~ spl24_70), inference(avatar_component_clause, [], [f977])).
fof(f456, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f9701, plain, (~ spl24_58 | ~ spl24_73), inference(avatar_split_clause, [], [f9698, f990, f927])).
fof(f927, plain, (spl24_58 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_58])])).
fof(f990, plain, (spl24_73 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_73])])).
fof(f9698, plain, (~ (e2 = op(e2, e3)) | ~ spl24_73), inference(backward_demodulation, [], [f454, f992])).
fof(f992, plain, ((e2 = op(e2, e0)) | ~ spl24_73), inference(avatar_component_clause, [], [f990])).
fof(f454, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f9693, plain, (~ spl24_43 | ~ spl24_79 | ~ spl24_95 | spl24_180), inference(avatar_split_clause, [], [f9691, f1570, f1082, f1015, f864])).
fof(f1015, plain, (spl24_79 <=> (e3 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_79])])).
fof(f1082, plain, (spl24_95 <=> (e4 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_95])])).
fof(f1570, plain, (spl24_180 <=> (e2 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_180])])).
fof(f9691, plain, (~ (e2 = op(e3, e1)) | (~ spl24_79 | ~ spl24_95 | spl24_180)), inference(backward_demodulation, [], [f9660, f1017])).
fof(f1017, plain, ((e3 = op(e1, e4)) | ~ spl24_79), inference(avatar_component_clause, [], [f1015])).
fof(f9660, plain, (~ (e2 = op(op(e1, e4), e1)) | (~ spl24_95 | spl24_180)), inference(backward_demodulation, [], [f1572, f1084])).
fof(f1084, plain, ((e4 = op(e1, e1)) | ~ spl24_95), inference(avatar_component_clause, [], [f1082])).
fof(f1572, plain, (~ (e2 = op(op(e1, op(e1, e1)), e1)) | spl24_180), inference(avatar_component_clause, [], [f1570])).
fof(f9686, plain, (~ spl24_58 | ~ spl24_83), inference(avatar_split_clause, [], [f9680, f1032, f927])).
fof(f1032, plain, (spl24_83 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_83])])).
fof(f9680, plain, (~ (e2 = op(e2, e3)) | ~ spl24_83), inference(backward_demodulation, [], [f416, f1034])).
fof(f1034, plain, ((e2 = op(e1, e3)) | ~ spl24_83), inference(avatar_component_clause, [], [f1032])).
fof(f416, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f9679, plain, (~ spl24_76 | ~ spl24_86), inference(avatar_split_clause, [], [f9673, f1045, f1003])).
fof(f1003, plain, (spl24_76 <=> (e0 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_76])])).
fof(f1045, plain, (spl24_86 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_86])])).
fof(f9673, plain, (~ (e0 = op(e1, e4)) | ~ spl24_86), inference(backward_demodulation, [], [f450, f1047])).
fof(f1047, plain, ((e0 = op(e1, e2)) | ~ spl24_86), inference(avatar_component_clause, [], [f1045])).
fof(f450, plain, ~ (op(e1, e2) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f9678, plain, (~ spl24_11 | ~ spl24_86), inference(avatar_split_clause, [], [f9671, f1045, f730])).
fof(f9671, plain, (~ (e0 = op(e4, e2)) | ~ spl24_86), inference(backward_demodulation, [], [f408, f1047])).
fof(f408, plain, ~ (op(e1, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f9668, plain, (~ spl24_76 | ~ spl24_95 | spl24_205), inference(avatar_split_clause, [], [f9663, f1719, f1082, f1003])).
fof(f1719, plain, (spl24_205 <=> (e0 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_205])])).
fof(f9663, plain, (~ (e0 = op(e1, e4)) | (~ spl24_95 | spl24_205)), inference(backward_demodulation, [], [f1721, f1084])).
fof(f1721, plain, (~ (e0 = op(e1, op(e1, e1))) | spl24_205), inference(avatar_component_clause, [], [f1719])).
fof(f9667, plain, (~ spl24_79 | ~ spl24_95 | spl24_178), inference(avatar_split_clause, [], [f9659, f1558, f1082, f1015])).
fof(f1558, plain, (spl24_178 <=> (e3 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_178])])).
fof(f9659, plain, (~ (e3 = op(e1, e4)) | (~ spl24_95 | spl24_178)), inference(backward_demodulation, [], [f1560, f1084])).
fof(f1560, plain, (~ (e3 = op(e1, op(e1, e1))) | spl24_178), inference(avatar_component_clause, [], [f1558])).
fof(f9665, plain, (~ spl24_90 | ~ spl24_95), inference(avatar_split_clause, [], [f9652, f1082, f1061])).
fof(f1061, plain, (spl24_90 <=> (e4 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_90])])).
fof(f9652, plain, (~ (e4 = op(e1, e2)) | ~ spl24_95), inference(backward_demodulation, [], [f446, f1084])).
fof(f446, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f9664, plain, (~ spl24_70 | ~ spl24_95), inference(avatar_split_clause, [], [f9650, f1082, f977])).
fof(f9650, plain, (~ (e4 = op(e2, e1)) | ~ spl24_95), inference(backward_demodulation, [], [f396, f1084])).
fof(f396, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f9649, plain, (~ spl24_82 | ~ spl24_97), inference(avatar_split_clause, [], [f9643, f1091, f1028])).
fof(f1028, plain, (spl24_82 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_82])])).
fof(f1091, plain, (spl24_97 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_97])])).
fof(f9643, plain, (~ (e1 = op(e1, e3)) | ~ spl24_97), inference(backward_demodulation, [], [f444, f1093])).
fof(f1093, plain, ((e1 = op(e1, e0)) | ~ spl24_97), inference(avatar_component_clause, [], [f1091])).
fof(f444, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f9648, plain, (~ spl24_22 | ~ spl24_97), inference(avatar_split_clause, [], [f9640, f1091, f776])).
fof(f776, plain, (spl24_22 <=> (e1 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_22])])).
fof(f9640, plain, (~ (e1 = op(e4, e0)) | ~ spl24_97), inference(backward_demodulation, [], [f388, f1093])).
fof(f388, plain, ~ (op(e1, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f9637, plain, (~ spl24_30 | ~ spl24_105), inference(avatar_split_clause, [], [f9636, f1124, f809])).
fof(f1124, plain, (spl24_105 <=> (e4 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_105])])).
fof(f9636, plain, (~ (e4 = op(e3, e4)) | ~ spl24_105), inference(backward_demodulation, [], [f424, f1126])).
fof(f1126, plain, ((e4 = op(e0, e4)) | ~ spl24_105), inference(avatar_component_clause, [], [f1124])).
fof(f424, plain, ~ (op(e0, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f9634, plain, (~ spl24_34 | ~ spl24_109), inference(avatar_split_clause, [], [f9630, f1141, f826])).
fof(f1141, plain, (spl24_109 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_109])])).
fof(f9630, plain, (~ (e3 = op(e3, e3)) | ~ spl24_109), inference(backward_demodulation, [], [f414, f1143])).
fof(f1143, plain, ((e3 = op(e0, e3)) | ~ spl24_109), inference(avatar_component_clause, [], [f1141])).
fof(f414, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f9627, plain, (~ spl24_38 | ~ spl24_113), inference(avatar_split_clause, [], [f9621, f1158, f843])).
fof(f1158, plain, (spl24_113 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_113])])).
fof(f9621, plain, (~ (e2 = op(e3, e2)) | ~ spl24_113), inference(backward_demodulation, [], [f404, f1160])).
fof(f1160, plain, ((e2 = op(e0, e2)) | ~ spl24_113), inference(avatar_component_clause, [], [f1158])).
fof(f404, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f9618, plain, (~ spl24_102 | ~ spl24_117), inference(avatar_split_clause, [], [f9613, f1175, f1112])).
fof(f1112, plain, (spl24_102 <=> (e1 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_102])])).
fof(f1175, plain, (spl24_117 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_117])])).
fof(f9613, plain, (~ (e1 = op(e0, e4)) | ~ spl24_117), inference(backward_demodulation, [], [f438, f1177])).
fof(f1177, plain, ((e1 = op(e0, e1)) | ~ spl24_117), inference(avatar_component_clause, [], [f1175])).
fof(f438, plain, ~ (op(e0, e1) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f9617, plain, (~ spl24_42 | ~ spl24_117), inference(avatar_split_clause, [], [f9610, f1175, f860])).
fof(f860, plain, (spl24_42 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_42])])).
fof(f9610, plain, (~ (e1 = op(e3, e1)) | ~ spl24_117), inference(backward_demodulation, [], [f394, f1177])).
fof(f394, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f9592, plain, (~ spl24_106 | ~ spl24_121), inference(avatar_split_clause, [], [f9575, f1192, f1129])).
fof(f1129, plain, (spl24_106 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_106])])).
fof(f1192, plain, (spl24_121 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_121])])).
fof(f9575, plain, (~ (e0 = op(e0, e3)) | ~ spl24_121), inference(backward_demodulation, [], [f434, f1194])).
fof(f1194, plain, ((e0 = op(e0, e0)) | ~ spl24_121), inference(avatar_component_clause, [], [f1192])).
fof(f434, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f9568, plain, (spl24_25 | ~ spl24_126), inference(avatar_split_clause, [], [f9554, f1213, f788])).
fof(f1213, plain, (spl24_126 <=> (e0 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl24_126])])).
fof(f9554, plain, ((e4 = op(e4, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f330, f1215])).
fof(f1215, plain, ((e0 = unit) | ~ spl24_126), inference(avatar_component_clause, [], [f1213])).
fof(f330, plain, (e4 = op(e4, unit)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)) & (e4 = op(e4, unit)) & (e4 = op(unit, e4)) & (e3 = op(e3, unit)) & (e3 = op(unit, e3)) & (e2 = op(e2, unit)) & (e2 = op(unit, e2)) & (e1 = op(e1, unit)) & (e1 = op(unit, e1)) & (e0 = op(e0, unit)) & (e0 = op(unit, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax2)).
fof(f9567, plain, (spl24_105 | ~ spl24_126), inference(avatar_split_clause, [], [f9553, f1213, f1124])).
fof(f9553, plain, ((e4 = op(e0, e4)) | ~ spl24_126), inference(backward_demodulation, [], [f329, f1215])).
fof(f329, plain, (e4 = op(unit, e4)), inference(cnf_transformation, [], [f2])).
fof(f9566, plain, (spl24_49 | ~ spl24_126), inference(avatar_split_clause, [], [f9552, f1213, f889])).
fof(f9552, plain, ((e3 = op(e3, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f328, f1215])).
fof(f328, plain, (e3 = op(e3, unit)), inference(cnf_transformation, [], [f2])).
fof(f9565, plain, (spl24_109 | ~ spl24_126), inference(avatar_split_clause, [], [f9551, f1213, f1141])).
fof(f9551, plain, ((e3 = op(e0, e3)) | ~ spl24_126), inference(backward_demodulation, [], [f327, f1215])).
fof(f327, plain, (e3 = op(unit, e3)), inference(cnf_transformation, [], [f2])).
fof(f9564, plain, (spl24_73 | ~ spl24_126), inference(avatar_split_clause, [], [f9550, f1213, f990])).
fof(f9550, plain, ((e2 = op(e2, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f326, f1215])).
fof(f326, plain, (e2 = op(e2, unit)), inference(cnf_transformation, [], [f2])).
fof(f9563, plain, (spl24_113 | ~ spl24_126), inference(avatar_split_clause, [], [f9549, f1213, f1158])).
fof(f9549, plain, ((e2 = op(e0, e2)) | ~ spl24_126), inference(backward_demodulation, [], [f325, f1215])).
fof(f325, plain, (e2 = op(unit, e2)), inference(cnf_transformation, [], [f2])).
fof(f9562, plain, (spl24_97 | ~ spl24_126), inference(avatar_split_clause, [], [f9548, f1213, f1091])).
fof(f9548, plain, ((e1 = op(e1, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f324, f1215])).
fof(f324, plain, (e1 = op(e1, unit)), inference(cnf_transformation, [], [f2])).
fof(f9561, plain, (spl24_117 | ~ spl24_126), inference(avatar_split_clause, [], [f9547, f1213, f1175])).
fof(f9547, plain, ((e1 = op(e0, e1)) | ~ spl24_126), inference(backward_demodulation, [], [f323, f1215])).
fof(f323, plain, (e1 = op(unit, e1)), inference(cnf_transformation, [], [f2])).
fof(f9560, plain, (spl24_121 | ~ spl24_126), inference(avatar_split_clause, [], [f9546, f1213, f1192])).
fof(f9546, plain, ((e0 = op(e0, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f322, f1215])).
fof(f322, plain, (e0 = op(e0, unit)), inference(cnf_transformation, [], [f2])).
fof(f9542, plain, (~ spl24_169 | ~ spl24_181), inference(avatar_contradiction_clause, [], [f9541])).
fof(f9541, plain, ($false | (~ spl24_169 | ~ spl24_181)), inference(subsumption_resolution, [], [f9540, f487])).
fof(f9540, plain, ((e1 = e3) | (~ spl24_169 | ~ spl24_181)), inference(forward_demodulation, [], [f1517, f1576])).
fof(f1576, plain, ((e1 = op(e2, op(e2, e2))) | ~ spl24_181), inference(avatar_component_clause, [], [f1575])).
fof(f1575, plain, (spl24_181 <=> (e1 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_181])])).
fof(f1517, plain, ((e3 = op(e2, op(e2, e2))) | ~ spl24_169), inference(avatar_component_clause, [], [f1516])).
fof(f1516, plain, (spl24_169 <=> (e3 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_169])])).
fof(f9533, plain, (~ spl24_86 | ~ spl24_181 | spl24_195), inference(avatar_split_clause, [], [f9530, f1652, f1575, f1045])).
fof(f1652, plain, (spl24_195 <=> (e0 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_195])])).
fof(f9530, plain, (~ (e0 = op(e1, e2)) | (~ spl24_181 | spl24_195)), inference(backward_demodulation, [], [f1654, f1576])).
fof(f1654, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | spl24_195), inference(avatar_component_clause, [], [f1652])).
fof(f9499, plain, (~ spl24_14 | ~ spl24_19), inference(avatar_split_clause, [], [f9475, f763, f742])).
fof(f9475, plain, (~ (e3 = op(e4, e2)) | ~ spl24_19), inference(backward_demodulation, [], [f476, f765])).
fof(f476, plain, ~ (op(e4, e1) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f9495, plain, (~ spl24_3 | ~ spl24_54 | spl24_207), inference(avatar_contradiction_clause, [], [f9494])).
fof(f9494, plain, ($false | (~ spl24_3 | ~ spl24_54 | spl24_207)), inference(subsumption_resolution, [], [f9493, f912])).
fof(f9493, plain, (~ (e3 = op(e2, e4)) | (~ spl24_3 | spl24_207)), inference(backward_demodulation, [], [f1737, f698])).
fof(f698, plain, ((e2 = op(e4, e4)) | ~ spl24_3), inference(avatar_component_clause, [], [f696])).
fof(f696, plain, (spl24_3 <=> (e2 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_3])])).
fof(f1737, plain, (~ (e3 = op(op(e4, e4), e4)) | spl24_207), inference(avatar_component_clause, [], [f1735])).
fof(f1735, plain, (spl24_207 <=> (e3 = op(op(e4, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_207])])).
fof(f9487, plain, (~ spl24_5 | ~ spl24_10), inference(avatar_split_clause, [], [f9486, f725, f704])).
fof(f704, plain, (spl24_5 <=> (e4 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_5])])).
fof(f9486, plain, (~ (e4 = op(e4, e4)) | ~ spl24_10), inference(backward_demodulation, [], [f481, f727])).
fof(f727, plain, ((e4 = op(e4, e3)) | ~ spl24_10), inference(avatar_component_clause, [], [f725])).
fof(f481, plain, ~ (op(e4, e3) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f9479, plain, (~ spl24_9 | ~ spl24_19), inference(avatar_split_clause, [], [f9476, f763, f721])).
fof(f721, plain, (spl24_9 <=> (e3 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_9])])).
fof(f9476, plain, (~ (e3 = op(e4, e3)) | ~ spl24_19), inference(backward_demodulation, [], [f477, f765])).
fof(f477, plain, ~ (op(e4, e1) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f9474, plain, (~ spl24_17 | ~ spl24_22), inference(avatar_split_clause, [], [f9467, f776, f755])).
fof(f755, plain, (spl24_17 <=> (e1 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_17])])).
fof(f9467, plain, (~ (e1 = op(e4, e1)) | ~ spl24_22), inference(backward_demodulation, [], [f472, f778])).
fof(f778, plain, ((e1 = op(e4, e0)) | ~ spl24_22), inference(avatar_component_clause, [], [f776])).
fof(f472, plain, ~ (op(e4, e0) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f9457, plain, (~ spl24_29 | ~ spl24_34), inference(avatar_split_clause, [], [f9445, f826, f805])).
fof(f805, plain, (spl24_29 <=> (e3 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_29])])).
fof(f9445, plain, (~ (e3 = op(e3, e4)) | ~ spl24_34), inference(backward_demodulation, [], [f471, f828])).
fof(f828, plain, ((e3 = op(e3, e3)) | ~ spl24_34), inference(avatar_component_clause, [], [f826])).
fof(f9456, plain, (~ spl24_9 | ~ spl24_34), inference(avatar_split_clause, [], [f9444, f826, f721])).
fof(f9444, plain, (~ (e3 = op(e4, e3)) | ~ spl24_34), inference(backward_demodulation, [], [f421, f828])).
fof(f9443, plain, (~ spl24_33 | ~ spl24_38), inference(avatar_split_clause, [], [f9438, f843, f822])).
fof(f822, plain, (spl24_33 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_33])])).
fof(f9438, plain, (~ (e2 = op(e3, e3)) | ~ spl24_38), inference(backward_demodulation, [], [f469, f845])).
fof(f845, plain, ((e2 = op(e3, e2)) | ~ spl24_38), inference(avatar_component_clause, [], [f843])).
fof(f469, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f9442, plain, (~ spl24_13 | ~ spl24_38), inference(avatar_split_clause, [], [f9437, f843, f738])).
fof(f738, plain, (spl24_13 <=> (e2 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_13])])).
fof(f9437, plain, (~ (e2 = op(e4, e2)) | ~ spl24_38), inference(backward_demodulation, [], [f411, f845])).
fof(f411, plain, ~ (op(e3, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f9436, plain, (~ spl24_17 | ~ spl24_42), inference(avatar_split_clause, [], [f9429, f860, f755])).
fof(f9429, plain, (~ (e1 = op(e4, e1)) | ~ spl24_42), inference(backward_demodulation, [], [f401, f862])).
fof(f862, plain, ((e1 = op(e3, e1)) | ~ spl24_42), inference(avatar_component_clause, [], [f860])).
fof(f401, plain, ~ (op(e3, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f9427, plain, (~ spl24_21 | ~ spl24_46), inference(avatar_split_clause, [], [f9418, f877, f772])).
fof(f772, plain, (spl24_21 <=> (e0 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_21])])).
fof(f877, plain, (spl24_46 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_46])])).
fof(f9418, plain, (~ (e0 = op(e4, e0)) | ~ spl24_46), inference(backward_demodulation, [], [f391, f879])).
fof(f879, plain, ((e0 = op(e3, e0)) | ~ spl24_46), inference(avatar_component_clause, [], [f877])).
fof(f391, plain, ~ (op(e3, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f9417, plain, (~ spl24_29 | ~ spl24_54), inference(avatar_split_clause, [], [f9414, f910, f805])).
fof(f9414, plain, (~ (e3 = op(e3, e4)) | ~ spl24_54), inference(backward_demodulation, [], [f429, f912])).
fof(f429, plain, ~ (op(e2, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f9413, plain, (~ spl24_53 | ~ spl24_58), inference(avatar_split_clause, [], [f9409, f927, f906])).
fof(f906, plain, (spl24_53 <=> (e2 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_53])])).
fof(f9409, plain, (~ (e2 = op(e2, e4)) | ~ spl24_58), inference(backward_demodulation, [], [f461, f929])).
fof(f929, plain, ((e2 = op(e2, e3)) | ~ spl24_58), inference(avatar_component_clause, [], [f927])).
fof(f461, plain, ~ (op(e2, e3) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f9406, plain, (~ spl24_60 | ~ spl24_75), inference(avatar_split_clause, [], [f9404, f998, f935])).
fof(f935, plain, (spl24_60 <=> (e4 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_60])])).
fof(f998, plain, (spl24_75 <=> (e4 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_75])])).
fof(f9404, plain, (~ (e4 = op(e2, e3)) | ~ spl24_75), inference(backward_demodulation, [], [f454, f1000])).
fof(f1000, plain, ((e4 = op(e2, e0)) | ~ spl24_75), inference(avatar_component_clause, [], [f998])).
fof(f9394, plain, (~ spl24_77 | ~ spl24_82), inference(avatar_split_clause, [], [f9390, f1028, f1007])).
fof(f1007, plain, (spl24_77 <=> (e1 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_77])])).
fof(f9390, plain, (~ (e1 = op(e1, e4)) | ~ spl24_82), inference(backward_demodulation, [], [f451, f1030])).
fof(f1030, plain, ((e1 = op(e1, e3)) | ~ spl24_82), inference(avatar_component_clause, [], [f1028])).
fof(f451, plain, ~ (op(e1, e3) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f9386, plain, (~ spl24_19 | ~ spl24_90 | ~ spl24_93 | spl24_174), inference(avatar_split_clause, [], [f9385, f1540, f1074, f1061, f763])).
fof(f1074, plain, (spl24_93 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_93])])).
fof(f1540, plain, (spl24_174 <=> (e3 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_174])])).
fof(f9385, plain, (~ (e3 = op(e4, e1)) | (~ spl24_90 | ~ spl24_93 | spl24_174)), inference(backward_demodulation, [], [f9295, f1063])).
fof(f1063, plain, ((e4 = op(e1, e2)) | ~ spl24_90), inference(avatar_component_clause, [], [f1061])).
fof(f9295, plain, (~ (e3 = op(op(e1, e2), e1)) | (~ spl24_93 | spl24_174)), inference(forward_demodulation, [], [f1542, f1076])).
fof(f1076, plain, ((e2 = op(e1, e1)) | ~ spl24_93), inference(avatar_component_clause, [], [f1074])).
fof(f1542, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | spl24_174), inference(avatar_component_clause, [], [f1540])).
fof(f9377, plain, (~ spl24_89 | ~ spl24_99), inference(avatar_split_clause, [], [f9372, f1099, f1057])).
fof(f1057, plain, (spl24_89 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_89])])).
fof(f1099, plain, (spl24_99 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_99])])).
fof(f9372, plain, (~ (e3 = op(e1, e2)) | ~ spl24_99), inference(backward_demodulation, [], [f443, f1101])).
fof(f1101, plain, ((e3 = op(e1, e0)) | ~ spl24_99), inference(avatar_component_clause, [], [f1099])).
fof(f443, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f9376, plain, (~ spl24_74 | ~ spl24_99), inference(avatar_split_clause, [], [f9369, f1099, f994])).
fof(f994, plain, (spl24_74 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_74])])).
fof(f9369, plain, (~ (e3 = op(e2, e0)) | ~ spl24_99), inference(backward_demodulation, [], [f386, f1101])).
fof(f386, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f4])).
fof(f9368, plain, (~ spl24_77 | ~ spl24_102), inference(avatar_split_clause, [], [f9361, f1112, f1007])).
fof(f9361, plain, (~ (e1 = op(e1, e4)) | ~ spl24_102), inference(backward_demodulation, [], [f422, f1114])).
fof(f1114, plain, ((e1 = op(e0, e4)) | ~ spl24_102), inference(avatar_component_clause, [], [f1112])).
fof(f422, plain, ~ (op(e0, e4) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f9360, plain, (~ spl24_101 | ~ spl24_106), inference(avatar_split_clause, [], [f9354, f1129, f1108])).
fof(f1108, plain, (spl24_101 <=> (e0 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_101])])).
fof(f9354, plain, (~ (e0 = op(e0, e4)) | ~ spl24_106), inference(backward_demodulation, [], [f441, f1131])).
fof(f1131, plain, ((e0 = op(e0, e3)) | ~ spl24_106), inference(avatar_component_clause, [], [f1129])).
fof(f441, plain, ~ (op(e0, e3) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f9359, plain, (~ spl24_81 | ~ spl24_106), inference(avatar_split_clause, [], [f9350, f1129, f1024])).
fof(f1024, plain, (spl24_81 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_81])])).
fof(f9350, plain, (~ (e0 = op(e1, e3)) | ~ spl24_106), inference(backward_demodulation, [], [f412, f1131])).
fof(f412, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f9346, plain, (~ spl24_89 | ~ spl24_114), inference(avatar_split_clause, [], [f9336, f1162, f1057])).
fof(f1162, plain, (spl24_114 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_114])])).
fof(f9336, plain, (~ (e3 = op(e1, e2)) | ~ spl24_114), inference(backward_demodulation, [], [f402, f1164])).
fof(f1164, plain, ((e3 = op(e0, e2)) | ~ spl24_114), inference(avatar_component_clause, [], [f1162])).
fof(f402, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f4])).
fof(f9335, plain, (~ spl24_115 | ~ spl24_120), inference(avatar_split_clause, [], [f9331, f1187, f1166])).
fof(f1166, plain, (spl24_115 <=> (e4 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_115])])).
fof(f1187, plain, (spl24_120 <=> (e4 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_120])])).
fof(f9331, plain, (~ (e4 = op(e0, e2)) | ~ spl24_120), inference(backward_demodulation, [], [f436, f1189])).
fof(f1189, plain, ((e4 = op(e0, e1)) | ~ spl24_120), inference(avatar_component_clause, [], [f1187])).
fof(f436, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f4])).
fof(f9334, plain, (~ spl24_45 | ~ spl24_120), inference(avatar_split_clause, [], [f9329, f1187, f872])).
fof(f872, plain, (spl24_45 <=> (e4 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_45])])).
fof(f9329, plain, (~ (e4 = op(e3, e1)) | ~ spl24_120), inference(backward_demodulation, [], [f394, f1189])).
fof(f9328, plain, (spl24_10 | ~ spl24_129), inference(avatar_split_clause, [], [f9317, f1225, f725])).
fof(f1225, plain, (spl24_129 <=> (e3 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl24_129])])).
fof(f9317, plain, ((e4 = op(e4, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f330, f1227])).
fof(f1227, plain, ((e3 = unit) | ~ spl24_129), inference(avatar_component_clause, [], [f1225])).
fof(f9327, plain, (spl24_30 | ~ spl24_129), inference(avatar_split_clause, [], [f9316, f1225, f809])).
fof(f9316, plain, ((e4 = op(e3, e4)) | ~ spl24_129), inference(backward_demodulation, [], [f329, f1227])).
fof(f9326, plain, (spl24_34 | ~ spl24_129), inference(avatar_split_clause, [], [f9315, f1225, f826])).
fof(f9315, plain, ((e3 = op(e3, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f328, f1227])).
fof(f9324, plain, (spl24_58 | ~ spl24_129), inference(avatar_split_clause, [], [f9313, f1225, f927])).
fof(f9313, plain, ((e2 = op(e2, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f326, f1227])).
fof(f9323, plain, (spl24_38 | ~ spl24_129), inference(avatar_split_clause, [], [f9312, f1225, f843])).
fof(f9312, plain, ((e2 = op(e3, e2)) | ~ spl24_129), inference(backward_demodulation, [], [f325, f1227])).
fof(f9322, plain, (spl24_82 | ~ spl24_129), inference(avatar_split_clause, [], [f9311, f1225, f1028])).
fof(f9311, plain, ((e1 = op(e1, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f324, f1227])).
fof(f9321, plain, (spl24_42 | ~ spl24_129), inference(avatar_split_clause, [], [f9310, f1225, f860])).
fof(f9310, plain, ((e1 = op(e3, e1)) | ~ spl24_129), inference(backward_demodulation, [], [f323, f1227])).
fof(f9320, plain, (spl24_106 | ~ spl24_129), inference(avatar_split_clause, [], [f9309, f1225, f1129])).
fof(f9309, plain, ((e0 = op(e0, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f322, f1227])).
fof(f9319, plain, (spl24_46 | ~ spl24_129), inference(avatar_split_clause, [], [f9308, f1225, f877])).
fof(f9308, plain, ((e0 = op(e3, e0)) | ~ spl24_129), inference(backward_demodulation, [], [f321, f1227])).
fof(f321, plain, (e0 = op(unit, e0)), inference(cnf_transformation, [], [f2])).
fof(f9301, plain, (~ spl24_115 | ~ spl24_62 | ~ spl24_66 | spl24_168), inference(avatar_split_clause, [], [f9300, f1512, f961, f944, f1166])).
fof(f944, plain, (spl24_62 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_62])])).
fof(f961, plain, (spl24_66 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_66])])).
fof(f1512, plain, (spl24_168 <=> (e4 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_168])])).
fof(f9300, plain, (~ (e4 = op(e0, e2)) | (~ spl24_62 | ~ spl24_66 | spl24_168)), inference(forward_demodulation, [], [f9299, f963])).
fof(f963, plain, ((e0 = op(e2, e1)) | ~ spl24_66), inference(avatar_component_clause, [], [f961])).
fof(f9299, plain, (~ (e4 = op(op(e2, e1), e2)) | (~ spl24_62 | spl24_168)), inference(forward_demodulation, [], [f1514, f946])).
fof(f946, plain, ((e1 = op(e2, e2)) | ~ spl24_62), inference(avatar_component_clause, [], [f944])).
fof(f1514, plain, (~ (e4 = op(op(e2, op(e2, e2)), e2)) | spl24_168), inference(avatar_component_clause, [], [f1512])).
fof(f9287, plain, (~ spl24_115 | ~ spl24_123 | spl24_191), inference(avatar_split_clause, [], [f9286, f1632, f1200, f1166])).
fof(f1200, plain, (spl24_123 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl24_123])])).
fof(f1632, plain, (spl24_191 <=> (e4 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_191])])).
fof(f9286, plain, (~ (e4 = op(e0, e2)) | (~ spl24_123 | spl24_191)), inference(forward_demodulation, [], [f1634, f1202])).
fof(f1202, plain, ((op(e0, e0) = e2) | ~ spl24_123), inference(avatar_component_clause, [], [f1200])).
fof(f1634, plain, (~ (e4 = op(e0, op(e0, e0))) | spl24_191), inference(avatar_component_clause, [], [f1632])).
fof(f9283, plain, (~ spl24_89 | ~ spl24_62 | spl24_208), inference(avatar_split_clause, [], [f9282, f1740, f944, f1057])).
fof(f9282, plain, (~ (e3 = op(e1, e2)) | (~ spl24_62 | spl24_208)), inference(forward_demodulation, [], [f1742, f946])).
fof(f1742, plain, (~ (e3 = op(op(e2, e2), e2)) | spl24_208), inference(avatar_component_clause, [], [f1740])).
fof(f9281, plain, (~ spl24_90 | ~ spl24_62 | spl24_212), inference(avatar_split_clause, [], [f9280, f1780, f944, f1061])).
fof(f1780, plain, (spl24_212 <=> (e4 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_212])])).
fof(f9280, plain, (~ (e4 = op(e1, e2)) | (~ spl24_62 | spl24_212)), inference(forward_demodulation, [], [f1782, f946])).
fof(f1782, plain, (~ (e4 = op(op(e2, e2), e2)) | spl24_212), inference(avatar_component_clause, [], [f1780])).
fof(f9269, plain, (~ spl24_114 | ~ spl24_123 | spl24_194), inference(avatar_split_clause, [], [f9065, f1646, f1200, f1162])).
fof(f1646, plain, (spl24_194 <=> (e3 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_194])])).
fof(f9065, plain, (~ (e3 = op(e0, e2)) | (~ spl24_123 | spl24_194)), inference(backward_demodulation, [], [f1648, f1202])).
fof(f1648, plain, (~ (e3 = op(e0, op(e0, e0))) | spl24_194), inference(avatar_component_clause, [], [f1646])).
fof(f9264, plain, (~ spl24_2 | ~ spl24_5 | spl24_172), inference(avatar_split_clause, [], [f9256, f1531, f704, f692])).
fof(f692, plain, (spl24_2 <=> (e1 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_2])])).
fof(f1531, plain, (spl24_172 <=> (e1 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_172])])).
fof(f9256, plain, (~ (e1 = op(e4, e4)) | (~ spl24_5 | spl24_172)), inference(backward_demodulation, [], [f1533, f706])).
fof(f706, plain, ((e4 = op(e4, e4)) | ~ spl24_5), inference(avatar_component_clause, [], [f704])).
fof(f1533, plain, (~ (e1 = op(e4, op(e4, e4))) | spl24_172), inference(avatar_component_clause, [], [f1531])).
fof(f9244, plain, (~ spl24_2 | ~ spl24_17), inference(avatar_split_clause, [], [f9240, f755, f692])).
fof(f9240, plain, (~ (e1 = op(e4, e4)) | ~ spl24_17), inference(backward_demodulation, [], [f478, f757])).
fof(f757, plain, ((e1 = op(e4, e1)) | ~ spl24_17), inference(avatar_component_clause, [], [f755])).
fof(f478, plain, ~ (op(e4, e1) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f9237, plain, (~ spl24_16 | ~ spl24_21), inference(avatar_split_clause, [], [f9229, f772, f751])).
fof(f9229, plain, (~ (e0 = op(e4, e1)) | ~ spl24_21), inference(backward_demodulation, [], [f472, f774])).
fof(f774, plain, ((e0 = op(e4, e0)) | ~ spl24_21), inference(avatar_component_clause, [], [f772])).
fof(f9226, plain, (~ spl24_33 | ~ spl24_60 | spl24_211), inference(avatar_contradiction_clause, [], [f9225])).
fof(f9225, plain, ($false | (~ spl24_33 | ~ spl24_60 | spl24_211)), inference(subsumption_resolution, [], [f9223, f937])).
fof(f937, plain, ((e4 = op(e2, e3)) | ~ spl24_60), inference(avatar_component_clause, [], [f935])).
fof(f9223, plain, (~ (e4 = op(e2, e3)) | (~ spl24_33 | spl24_211)), inference(backward_demodulation, [], [f1777, f824])).
fof(f824, plain, ((e2 = op(e3, e3)) | ~ spl24_33), inference(avatar_component_clause, [], [f822])).
fof(f1777, plain, (~ (e4 = op(op(e3, e3), e3)) | spl24_211), inference(avatar_component_clause, [], [f1775])).
fof(f1775, plain, (spl24_211 <=> (e4 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_211])])).
fof(f9224, plain, (~ spl24_8 | ~ spl24_33), inference(avatar_split_clause, [], [f9216, f822, f717])).
fof(f717, plain, (spl24_8 <=> (e2 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_8])])).
fof(f9216, plain, (~ (e2 = op(e4, e3)) | ~ spl24_33), inference(backward_demodulation, [], [f421, f824])).
fof(f9215, plain, (~ spl24_26 | ~ spl24_36), inference(avatar_split_clause, [], [f9210, f835, f793])).
fof(f793, plain, (spl24_26 <=> (e0 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_26])])).
fof(f835, plain, (spl24_36 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_36])])).
fof(f9210, plain, (~ (e0 = op(e3, e4)) | ~ spl24_36), inference(backward_demodulation, [], [f470, f837])).
fof(f837, plain, ((e0 = op(e3, e2)) | ~ spl24_36), inference(avatar_component_clause, [], [f835])).
fof(f470, plain, ~ (op(e3, e2) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f9207, plain, (~ spl24_40 | ~ spl24_45), inference(avatar_split_clause, [], [f9204, f872, f851])).
fof(f851, plain, (spl24_40 <=> (e4 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_40])])).
fof(f9204, plain, (~ (e4 = op(e3, e2)) | ~ spl24_45), inference(backward_demodulation, [], [f466, f874])).
fof(f874, plain, ((e4 = op(e3, e1)) | ~ spl24_45), inference(avatar_component_clause, [], [f872])).
fof(f9202, plain, (~ spl24_32 | ~ spl24_47), inference(avatar_split_clause, [], [f9197, f881, f818])).
fof(f818, plain, (spl24_32 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_32])])).
fof(f881, plain, (spl24_47 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_47])])).
fof(f9197, plain, (~ (e1 = op(e3, e3)) | ~ spl24_47), inference(backward_demodulation, [], [f464, f883])).
fof(f883, plain, ((e1 = op(e3, e0)) | ~ spl24_47), inference(avatar_component_clause, [], [f881])).
fof(f9186, plain, (~ spl24_56 | ~ spl24_66), inference(avatar_split_clause, [], [f9178, f961, f919])).
fof(f919, plain, (spl24_56 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_56])])).
fof(f9178, plain, (~ (e0 = op(e2, e3)) | ~ spl24_66), inference(backward_demodulation, [], [f457, f963])).
fof(f457, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f9185, plain, (~ spl24_16 | ~ spl24_66), inference(avatar_split_clause, [], [f9177, f961, f751])).
fof(f9177, plain, (~ (e0 = op(e4, e1)) | ~ spl24_66), inference(backward_demodulation, [], [f400, f963])).
fof(f400, plain, ~ (op(e2, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f9175, plain, (~ spl24_54 | ~ spl24_74), inference(avatar_split_clause, [], [f9172, f994, f910])).
fof(f9172, plain, (~ (e3 = op(e2, e4)) | ~ spl24_74), inference(backward_demodulation, [], [f455, f996])).
fof(f996, plain, ((e3 = op(e2, e0)) | ~ spl24_74), inference(avatar_component_clause, [], [f994])).
fof(f455, plain, ~ (op(e2, e0) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f9160, plain, (~ spl24_56 | ~ spl24_81), inference(avatar_split_clause, [], [f9152, f1024, f919])).
fof(f9152, plain, (~ (e0 = op(e2, e3)) | ~ spl24_81), inference(backward_demodulation, [], [f416, f1026])).
fof(f1026, plain, ((e0 = op(e1, e3)) | ~ spl24_81), inference(avatar_component_clause, [], [f1024])).
fof(f9149, plain, (~ spl24_14 | ~ spl24_89), inference(avatar_split_clause, [], [f9142, f1057, f742])).
fof(f9142, plain, (~ (e3 = op(e4, e2)) | ~ spl24_89), inference(backward_demodulation, [], [f408, f1059])).
fof(f1059, plain, ((e3 = op(e1, e2)) | ~ spl24_89), inference(avatar_component_clause, [], [f1057])).
fof(f9140, plain, (~ spl24_70 | ~ spl24_93 | spl24_213), inference(avatar_split_clause, [], [f9134, f1786, f1074, f977])).
fof(f1786, plain, (spl24_213 <=> (e4 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_213])])).
fof(f9134, plain, (~ (e4 = op(e2, e1)) | (~ spl24_93 | spl24_213)), inference(backward_demodulation, [], [f1788, f1076])).
fof(f1788, plain, (~ (e4 = op(op(e1, e1), e1)) | spl24_213), inference(avatar_component_clause, [], [f1786])).
fof(f9139, plain, (~ spl24_86 | ~ spl24_93 | spl24_205), inference(avatar_split_clause, [], [f9133, f1719, f1074, f1045])).
fof(f9133, plain, (~ (e0 = op(e1, e2)) | (~ spl24_93 | spl24_205)), inference(backward_demodulation, [], [f1721, f1076])).
fof(f9137, plain, (~ spl24_66 | ~ spl24_93 | spl24_173), inference(avatar_split_clause, [], [f9126, f1536, f1074, f961])).
fof(f1536, plain, (spl24_173 <=> (e0 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_173])])).
fof(f9126, plain, (~ (e0 = op(e2, e1)) | (~ spl24_93 | spl24_173)), inference(backward_demodulation, [], [f1538, f1076])).
fof(f1538, plain, (~ (e0 = op(op(e1, e1), e1)) | spl24_173), inference(avatar_component_clause, [], [f1536])).
fof(f9136, plain, (~ spl24_78 | ~ spl24_93), inference(avatar_split_clause, [], [f9123, f1074, f1011])).
fof(f1011, plain, (spl24_78 <=> (e2 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_78])])).
fof(f9123, plain, (~ (e2 = op(e1, e4)) | ~ spl24_93), inference(backward_demodulation, [], [f448, f1076])).
fof(f448, plain, ~ (op(e1, e1) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f9135, plain, (~ spl24_43 | ~ spl24_93), inference(avatar_split_clause, [], [f9119, f1074, f864])).
fof(f9119, plain, (~ (e2 = op(e3, e1)) | ~ spl24_93), inference(backward_demodulation, [], [f397, f1076])).
fof(f397, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f9117, plain, (~ spl24_85 | ~ spl24_100), inference(avatar_split_clause, [], [f9114, f1103, f1040])).
fof(f1040, plain, (spl24_85 <=> (e4 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_85])])).
fof(f1103, plain, (spl24_100 <=> (e4 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_100])])).
fof(f9114, plain, (~ (e4 = op(e1, e3)) | ~ spl24_100), inference(backward_demodulation, [], [f444, f1105])).
fof(f1105, plain, ((e4 = op(e1, e0)) | ~ spl24_100), inference(avatar_component_clause, [], [f1103])).
fof(f9108, plain, (~ spl24_26 | ~ spl24_101), inference(avatar_split_clause, [], [f9102, f1108, f793])).
fof(f9102, plain, (~ (e0 = op(e3, e4)) | ~ spl24_101), inference(backward_demodulation, [], [f424, f1110])).
fof(f1110, plain, ((e0 = op(e0, e4)) | ~ spl24_101), inference(avatar_component_clause, [], [f1108])).
fof(f9099, plain, (~ spl24_32 | ~ spl24_107), inference(avatar_split_clause, [], [f9093, f1133, f818])).
fof(f1133, plain, (spl24_107 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_107])])).
fof(f9093, plain, (~ (e1 = op(e3, e3)) | ~ spl24_107), inference(backward_demodulation, [], [f414, f1135])).
fof(f1135, plain, ((e1 = op(e0, e3)) | ~ spl24_107), inference(avatar_component_clause, [], [f1133])).
fof(f9090, plain, (~ spl24_105 | ~ spl24_115), inference(avatar_split_clause, [], [f9088, f1166, f1124])).
fof(f9088, plain, (~ (e4 = op(e0, e4)) | ~ spl24_115), inference(backward_demodulation, [], [f440, f1168])).
fof(f1168, plain, ((e4 = op(e0, e2)) | ~ spl24_115), inference(avatar_component_clause, [], [f1166])).
fof(f440, plain, ~ (op(e0, e2) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f9089, plain, (~ spl24_40 | ~ spl24_115), inference(avatar_split_clause, [], [f9085, f1166, f851])).
fof(f9085, plain, (~ (e4 = op(e3, e2)) | ~ spl24_115), inference(backward_demodulation, [], [f404, f1168])).
fof(f9083, plain, (~ spl24_109 | ~ spl24_119), inference(avatar_split_clause, [], [f9079, f1183, f1141])).
fof(f1183, plain, (spl24_119 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_119])])).
fof(f9079, plain, (~ (e3 = op(e0, e3)) | ~ spl24_119), inference(backward_demodulation, [], [f437, f1185])).
fof(f1185, plain, ((e3 = op(e0, e1)) | ~ spl24_119), inference(avatar_component_clause, [], [f1183])).
fof(f437, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f9082, plain, (~ spl24_94 | ~ spl24_119), inference(avatar_split_clause, [], [f9074, f1183, f1078])).
fof(f1078, plain, (spl24_94 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_94])])).
fof(f9074, plain, (~ (e3 = op(e1, e1)) | ~ spl24_119), inference(backward_demodulation, [], [f392, f1185])).
fof(f392, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f4])).
fof(f9073, plain, (~ spl24_74 | ~ spl24_123 | spl24_210), inference(avatar_split_clause, [], [f9068, f1756, f1200, f994])).
fof(f1756, plain, (spl24_210 <=> (e3 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_210])])).
fof(f9068, plain, (~ (e3 = op(e2, e0)) | (~ spl24_123 | spl24_210)), inference(backward_demodulation, [], [f1758, f1202])).
fof(f1758, plain, (~ (e3 = op(op(e0, e0), e0)) | spl24_210), inference(avatar_component_clause, [], [f1756])).
fof(f9071, plain, (spl24_115 | ~ spl24_123 | ~ spl24_191), inference(avatar_split_clause, [], [f9064, f1632, f1200, f1166])).
fof(f9064, plain, ((e4 = op(e0, e2)) | (~ spl24_123 | ~ spl24_191)), inference(backward_demodulation, [], [f1633, f1202])).
fof(f1633, plain, ((e4 = op(e0, op(e0, e0))) | ~ spl24_191), inference(avatar_component_clause, [], [f1632])).
fof(f9070, plain, (~ spl24_113 | ~ spl24_123), inference(avatar_split_clause, [], [f9058, f1200, f1158])).
fof(f9058, plain, (~ (e2 = op(e0, e2)) | ~ spl24_123), inference(backward_demodulation, [], [f433, f1202])).
fof(f433, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f4])).
fof(f9069, plain, (~ spl24_73 | ~ spl24_123), inference(avatar_split_clause, [], [f9054, f1200, f990])).
fof(f9054, plain, (~ (e2 = op(e2, e0)) | ~ spl24_123), inference(backward_demodulation, [], [f383, f1202])).
fof(f383, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f4])).
fof(f9052, plain, (spl24_5 | ~ spl24_130), inference(avatar_split_clause, [], [f9042, f1229, f704])).
fof(f1229, plain, (spl24_130 <=> (e4 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl24_130])])).
fof(f9042, plain, ((e4 = op(e4, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f330, f1231])).
fof(f1231, plain, ((e4 = unit) | ~ spl24_130), inference(avatar_component_clause, [], [f1229])).
fof(f9050, plain, (spl24_29 | ~ spl24_130), inference(avatar_split_clause, [], [f9040, f1229, f805])).
fof(f9040, plain, ((e3 = op(e3, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f328, f1231])).
fof(f9049, plain, (spl24_9 | ~ spl24_130), inference(avatar_split_clause, [], [f9039, f1229, f721])).
fof(f9039, plain, ((e3 = op(e4, e3)) | ~ spl24_130), inference(backward_demodulation, [], [f327, f1231])).
fof(f9048, plain, (spl24_53 | ~ spl24_130), inference(avatar_split_clause, [], [f9038, f1229, f906])).
fof(f9038, plain, ((e2 = op(e2, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f326, f1231])).
fof(f9047, plain, (spl24_13 | ~ spl24_130), inference(avatar_split_clause, [], [f9037, f1229, f738])).
fof(f9037, plain, ((e2 = op(e4, e2)) | ~ spl24_130), inference(backward_demodulation, [], [f325, f1231])).
fof(f9046, plain, (spl24_77 | ~ spl24_130), inference(avatar_split_clause, [], [f9036, f1229, f1007])).
fof(f9036, plain, ((e1 = op(e1, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f324, f1231])).
fof(f9045, plain, (spl24_17 | ~ spl24_130), inference(avatar_split_clause, [], [f9035, f1229, f755])).
fof(f9035, plain, ((e1 = op(e4, e1)) | ~ spl24_130), inference(backward_demodulation, [], [f323, f1231])).
fof(f9044, plain, (spl24_101 | ~ spl24_130), inference(avatar_split_clause, [], [f9034, f1229, f1108])).
fof(f9034, plain, ((e0 = op(e0, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f322, f1231])).
fof(f9043, plain, (spl24_21 | ~ spl24_130), inference(avatar_split_clause, [], [f9033, f1229, f772])).
fof(f9033, plain, ((e0 = op(e4, e0)) | ~ spl24_130), inference(backward_demodulation, [], [f321, f1231])).
fof(f9021, plain, (spl24_89 | ~ spl24_62 | ~ spl24_208), inference(avatar_split_clause, [], [f9020, f1740, f944, f1057])).
fof(f9020, plain, ((e3 = op(e1, e2)) | (~ spl24_62 | ~ spl24_208)), inference(forward_demodulation, [], [f1741, f946])).
fof(f9017, plain, (~ spl24_86 | ~ spl24_62 | spl24_163), inference(avatar_split_clause, [], [f8898, f1490, f944, f1045])).
fof(f8898, plain, (~ (e0 = op(e1, e2)) | (~ spl24_62 | spl24_163)), inference(backward_demodulation, [], [f1492, f946])).
fof(f9012, plain, (~ spl24_66 | ~ spl24_62 | spl24_197), inference(avatar_split_clause, [], [f8903, f1663, f944, f961])).
fof(f1663, plain, (spl24_197 <=> (e0 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_197])])).
fof(f8903, plain, (~ (e0 = op(e2, e1)) | (~ spl24_62 | spl24_197)), inference(backward_demodulation, [], [f1665, f946])).
fof(f1665, plain, (~ (e0 = op(e2, op(e2, e2))) | spl24_197), inference(avatar_component_clause, [], [f1663])).
fof(f8986, plain, (~ spl24_1 | ~ spl24_16), inference(avatar_split_clause, [], [f8981, f751, f688])).
fof(f688, plain, (spl24_1 <=> (e0 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_1])])).
fof(f8981, plain, (~ (e0 = op(e4, e4)) | ~ spl24_16), inference(backward_demodulation, [], [f478, f753])).
fof(f8973, plain, (~ spl24_1 | ~ spl24_26), inference(avatar_split_clause, [], [f8968, f793, f688])).
fof(f8968, plain, (~ (e0 = op(e4, e4)) | ~ spl24_26), inference(backward_demodulation, [], [f431, f795])).
fof(f795, plain, ((e0 = op(e3, e4)) | ~ spl24_26), inference(avatar_component_clause, [], [f793])).
fof(f431, plain, ~ (op(e3, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f8967, plain, (~ spl24_32 | ~ spl24_43 | ~ spl24_56 | spl24_187), inference(avatar_contradiction_clause, [], [f8966])).
fof(f8966, plain, ($false | (~ spl24_32 | ~ spl24_43 | ~ spl24_56 | spl24_187)), inference(subsumption_resolution, [], [f8965, f921])).
fof(f921, plain, ((e0 = op(e2, e3)) | ~ spl24_56), inference(avatar_component_clause, [], [f919])).
fof(f8965, plain, (~ (e0 = op(e2, e3)) | (~ spl24_32 | ~ spl24_43 | spl24_187)), inference(forward_demodulation, [], [f8955, f866])).
fof(f8955, plain, (~ (e0 = op(op(e3, e1), e3)) | (~ spl24_32 | spl24_187)), inference(backward_demodulation, [], [f1616, f820])).
fof(f820, plain, ((e1 = op(e3, e3)) | ~ spl24_32), inference(avatar_component_clause, [], [f818])).
fof(f1616, plain, (~ (e0 = op(op(e3, op(e3, e3)), e3)) | spl24_187), inference(avatar_component_clause, [], [f1614])).
fof(f1614, plain, (spl24_187 <=> (e0 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_187])])).
fof(f8962, plain, (~ spl24_32 | ~ spl24_43 | spl24_167), inference(avatar_contradiction_clause, [], [f8961])).
fof(f8961, plain, ($false | (~ spl24_32 | ~ spl24_43 | spl24_167)), inference(subsumption_resolution, [], [f8951, f866])).
fof(f8951, plain, (~ (e2 = op(e3, e1)) | (~ spl24_32 | spl24_167)), inference(backward_demodulation, [], [f1509, f820])).
fof(f1509, plain, (~ (e2 = op(e3, op(e3, e3))) | spl24_167), inference(avatar_component_clause, [], [f1507])).
fof(f1507, plain, (spl24_167 <=> (e2 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_167])])).
fof(f8942, plain, (~ spl24_30 | ~ spl24_40), inference(avatar_split_clause, [], [f8941, f851, f809])).
fof(f8941, plain, (~ (e4 = op(e3, e4)) | ~ spl24_40), inference(backward_demodulation, [], [f470, f853])).
fof(f853, plain, ((e4 = op(e3, e2)) | ~ spl24_40), inference(avatar_component_clause, [], [f851])).
fof(f8907, plain, (~ spl24_52 | ~ spl24_62), inference(avatar_split_clause, [], [f8894, f944, f902])).
fof(f8894, plain, (~ (e1 = op(e2, e4)) | ~ spl24_62), inference(backward_demodulation, [], [f460, f946])).
fof(f460, plain, ~ (op(e2, e2) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f8906, plain, (~ spl24_12 | ~ spl24_62), inference(avatar_split_clause, [], [f8892, f944, f734])).
fof(f734, plain, (spl24_12 <=> (e1 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_12])])).
fof(f8892, plain, (~ (e1 = op(e4, e2)) | ~ spl24_62), inference(backward_demodulation, [], [f410, f946])).
fof(f410, plain, ~ (op(e2, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f8881, plain, (~ spl24_10 | ~ spl24_85), inference(avatar_split_clause, [], [f8880, f1040, f725])).
fof(f8880, plain, (~ (e4 = op(e4, e3)) | ~ spl24_85), inference(backward_demodulation, [], [f418, f1042])).
fof(f1042, plain, ((e4 = op(e1, e3)) | ~ spl24_85), inference(avatar_component_clause, [], [f1040])).
fof(f418, plain, ~ (op(e1, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f8877, plain, (~ spl24_61 | ~ spl24_86), inference(avatar_split_clause, [], [f8869, f1045, f940])).
fof(f940, plain, (spl24_61 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_61])])).
fof(f8869, plain, (~ (e0 = op(e2, e2)) | ~ spl24_86), inference(backward_demodulation, [], [f406, f1047])).
fof(f406, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f8868, plain, (spl24_43 | ~ spl24_94 | ~ spl24_201), inference(avatar_split_clause, [], [f8863, f1690, f1078, f864])).
fof(f1690, plain, (spl24_201 <=> (e2 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_201])])).
fof(f8863, plain, ((e2 = op(e3, e1)) | (~ spl24_94 | ~ spl24_201)), inference(backward_demodulation, [], [f1691, f1080])).
fof(f1080, plain, ((e3 = op(e1, e1)) | ~ spl24_94), inference(avatar_component_clause, [], [f1078])).
fof(f1691, plain, ((e2 = op(op(e1, e1), e1)) | ~ spl24_201), inference(avatar_component_clause, [], [f1690])).
fof(f8867, plain, (spl24_85 | ~ spl24_94 | ~ spl24_175), inference(avatar_split_clause, [], [f8860, f1544, f1078, f1040])).
fof(f1544, plain, (spl24_175 <=> (e4 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_175])])).
fof(f8860, plain, ((e4 = op(e1, e3)) | (~ spl24_94 | ~ spl24_175)), inference(backward_demodulation, [], [f1545, f1080])).
fof(f1545, plain, ((e4 = op(e1, op(e1, e1))) | ~ spl24_175), inference(avatar_component_clause, [], [f1544])).
fof(f8866, plain, (~ spl24_89 | ~ spl24_94), inference(avatar_split_clause, [], [f8856, f1078, f1057])).
fof(f8856, plain, (~ (e3 = op(e1, e2)) | ~ spl24_94), inference(backward_demodulation, [], [f446, f1080])).
fof(f8865, plain, (~ spl24_19 | ~ spl24_94), inference(avatar_split_clause, [], [f8855, f1078, f763])).
fof(f8855, plain, (~ (e3 = op(e4, e1)) | ~ spl24_94), inference(backward_demodulation, [], [f398, f1080])).
fof(f398, plain, ~ (op(e1, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f8755, plain, (~ spl24_14 | spl24_164 | ~ spl24_165), inference(avatar_split_clause, [], [f8753, f1498, f1494, f742])).
fof(f1494, plain, (spl24_164 <=> (e3 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_164])])).
fof(f1498, plain, (spl24_165 <=> (e4 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl24_165])])).
fof(f8753, plain, (~ (e3 = op(e4, e2)) | (spl24_164 | ~ spl24_165)), inference(backward_demodulation, [], [f1496, f1499])).
fof(f1499, plain, ((e4 = op(e2, op(e2, e2))) | ~ spl24_165), inference(avatar_component_clause, [], [f1498])).
fof(f1496, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | spl24_164), inference(avatar_component_clause, [], [f1494])).
fof(f8750, plain, (~ spl24_16 | ~ spl24_175 | spl24_203), inference(avatar_split_clause, [], [f8747, f1708, f1544, f751])).
fof(f1708, plain, (spl24_203 <=> (e0 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_203])])).
fof(f8747, plain, (~ (e0 = op(e4, e1)) | (~ spl24_175 | spl24_203)), inference(backward_demodulation, [], [f1710, f1545])).
fof(f1710, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | spl24_203), inference(avatar_component_clause, [], [f1708])).
fof(f8732, plain, (~ spl24_1 | ~ spl24_23 | ~ spl24_52 | spl24_170), inference(avatar_contradiction_clause, [], [f8731])).
fof(f8731, plain, ($false | (~ spl24_1 | ~ spl24_23 | ~ spl24_52 | spl24_170)), inference(subsumption_resolution, [], [f8728, f904])).
fof(f8728, plain, (~ (e1 = op(e2, e4)) | (~ spl24_1 | ~ spl24_23 | spl24_170)), inference(backward_demodulation, [], [f8547, f782])).
fof(f782, plain, ((e2 = op(e4, e0)) | ~ spl24_23), inference(avatar_component_clause, [], [f780])).
fof(f780, plain, (spl24_23 <=> (e2 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_23])])).
fof(f8547, plain, (~ (e1 = op(op(e4, e0), e4)) | (~ spl24_1 | spl24_170)), inference(forward_demodulation, [], [f1523, f690])).
fof(f690, plain, ((e0 = op(e4, e4)) | ~ spl24_1), inference(avatar_component_clause, [], [f688])).
fof(f1523, plain, (~ (e1 = op(op(e4, op(e4, e4)), e4)) | spl24_170), inference(avatar_component_clause, [], [f1521])).
fof(f1521, plain, (spl24_170 <=> (e1 = op(op(e4, op(e4, e4)), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_170])])).
fof(f8730, plain, (~ spl24_18 | ~ spl24_23), inference(avatar_split_clause, [], [f8723, f780, f759])).
fof(f759, plain, (spl24_18 <=> (e2 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_18])])).
fof(f8723, plain, (~ (e2 = op(e4, e1)) | ~ spl24_23), inference(backward_demodulation, [], [f472, f782])).
fof(f8700, plain, (~ spl24_27 | ~ spl24_42), inference(avatar_split_clause, [], [f8696, f860, f797])).
fof(f797, plain, (spl24_27 <=> (e1 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_27])])).
fof(f8696, plain, (~ (e1 = op(e3, e4)) | ~ spl24_42), inference(backward_demodulation, [], [f468, f862])).
fof(f468, plain, ~ (op(e3, e1) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f8682, plain, (~ spl24_27 | ~ spl24_52), inference(avatar_split_clause, [], [f8678, f902, f797])).
fof(f8678, plain, (~ (e1 = op(e3, e4)) | ~ spl24_52), inference(backward_demodulation, [], [f429, f904])).
fof(f8657, plain, (~ spl24_55 | ~ spl24_70), inference(avatar_split_clause, [], [f8656, f977, f914])).
fof(f914, plain, (spl24_55 <=> (e4 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_55])])).
fof(f8656, plain, (~ (e4 = op(e2, e4)) | ~ spl24_70), inference(backward_demodulation, [], [f458, f979])).
fof(f458, plain, ~ (op(e2, e1) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f8651, plain, (~ spl24_59 | ~ spl24_74), inference(avatar_split_clause, [], [f8647, f994, f931])).
fof(f931, plain, (spl24_59 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_59])])).
fof(f8647, plain, (~ (e3 = op(e2, e3)) | ~ spl24_74), inference(backward_demodulation, [], [f454, f996])).
fof(f8650, plain, (~ spl24_24 | ~ spl24_74), inference(avatar_split_clause, [], [f8644, f994, f784])).
fof(f784, plain, (spl24_24 <=> (e3 = op(e4, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_24])])).
fof(f8644, plain, (~ (e3 = op(e4, e0)) | ~ spl24_74), inference(backward_demodulation, [], [f390, f996])).
fof(f390, plain, ~ (op(e2, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f8635, plain, (~ spl24_39 | ~ spl24_89), inference(avatar_split_clause, [], [f8631, f1057, f847])).
fof(f847, plain, (spl24_39 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_39])])).
fof(f8631, plain, (~ (e3 = op(e3, e2)) | ~ spl24_89), inference(backward_demodulation, [], [f407, f1059])).
fof(f407, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f8629, plain, (~ spl24_18 | ~ spl24_91 | ~ spl24_100 | spl24_180), inference(avatar_split_clause, [], [f8628, f1570, f1103, f1066, f759])).
fof(f1066, plain, (spl24_91 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_91])])).
fof(f8628, plain, (~ (e2 = op(e4, e1)) | (~ spl24_91 | ~ spl24_100 | spl24_180)), inference(forward_demodulation, [], [f8618, f1105])).
fof(f8618, plain, (~ (e2 = op(op(e1, e0), e1)) | (~ spl24_91 | spl24_180)), inference(backward_demodulation, [], [f1572, f1068])).
fof(f1068, plain, ((e0 = op(e1, e1)) | ~ spl24_91), inference(avatar_component_clause, [], [f1066])).
fof(f8627, plain, (~ spl24_91 | ~ spl24_100 | spl24_175), inference(avatar_contradiction_clause, [], [f8626])).
fof(f8626, plain, ($false | (~ spl24_91 | ~ spl24_100 | spl24_175)), inference(subsumption_resolution, [], [f8617, f1105])).
fof(f8617, plain, (~ (e4 = op(e1, e0)) | (~ spl24_91 | spl24_175)), inference(backward_demodulation, [], [f1546, f1068])).
fof(f1546, plain, (~ (e4 = op(e1, op(e1, e1))) | spl24_175), inference(avatar_component_clause, [], [f1544])).
fof(f8623, plain, (~ spl24_81 | ~ spl24_91), inference(avatar_split_clause, [], [f8611, f1066, f1024])).
fof(f8611, plain, (~ (e0 = op(e1, e3)) | ~ spl24_91), inference(backward_demodulation, [], [f447, f1068])).
fof(f447, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f8622, plain, (~ spl24_41 | ~ spl24_91), inference(avatar_split_clause, [], [f8608, f1066, f856])).
fof(f856, plain, (spl24_41 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_41])])).
fof(f8608, plain, (~ (e0 = op(e3, e1)) | ~ spl24_91), inference(backward_demodulation, [], [f397, f1068])).
fof(f8606, plain, (~ spl24_95 | ~ spl24_100), inference(avatar_split_clause, [], [f8603, f1103, f1082])).
fof(f8603, plain, (~ (e4 = op(e1, e1)) | ~ spl24_100), inference(backward_demodulation, [], [f442, f1105])).
fof(f442, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f4])).
fof(f8590, plain, (~ spl24_15 | ~ spl24_115), inference(avatar_split_clause, [], [f8588, f1166, f746])).
fof(f746, plain, (spl24_15 <=> (e4 = op(e4, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_15])])).
fof(f8588, plain, (~ (e4 = op(e4, e2)) | ~ spl24_115), inference(backward_demodulation, [], [f405, f1168])).
fof(f405, plain, ~ (op(e0, e2) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f8584, plain, (~ spl24_108 | ~ spl24_118), inference(avatar_split_clause, [], [f8580, f1179, f1137])).
fof(f1137, plain, (spl24_108 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_108])])).
fof(f1179, plain, (spl24_118 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_118])])).
fof(f8580, plain, (~ (e2 = op(e0, e3)) | ~ spl24_118), inference(backward_demodulation, [], [f437, f1181])).
fof(f1181, plain, ((e2 = op(e0, e1)) | ~ spl24_118), inference(avatar_component_clause, [], [f1179])).
fof(f8583, plain, (~ spl24_18 | ~ spl24_118), inference(avatar_split_clause, [], [f8578, f1179, f759])).
fof(f8578, plain, (~ (e2 = op(e4, e1)) | ~ spl24_118), inference(backward_demodulation, [], [f395, f1181])).
fof(f395, plain, ~ (op(e0, e1) = op(e4, e1)), inference(cnf_transformation, [], [f4])).
fof(f8544, plain, (~ spl24_100 | ~ spl24_122 | spl24_214), inference(avatar_split_clause, [], [f8543, f1796, f1196, f1103])).
fof(f1196, plain, (spl24_122 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl24_122])])).
fof(f1796, plain, (spl24_214 <=> (e4 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_214])])).
fof(f8543, plain, (~ (e4 = op(e1, e0)) | (~ spl24_122 | spl24_214)), inference(forward_demodulation, [], [f1798, f1198])).
fof(f1198, plain, ((op(e0, e0) = e1) | ~ spl24_122), inference(avatar_component_clause, [], [f1196])).
fof(f1798, plain, (~ (e4 = op(op(e0, e0), e0)) | spl24_214), inference(avatar_component_clause, [], [f1796])).
fof(f8535, plain, (spl24_119 | ~ spl24_122 | ~ spl24_194), inference(avatar_split_clause, [], [f8534, f1646, f1196, f1183])).
fof(f8534, plain, ((e3 = op(e0, e1)) | (~ spl24_122 | ~ spl24_194)), inference(forward_demodulation, [], [f1647, f1198])).
fof(f1647, plain, ((e3 = op(e0, op(e0, e0))) | ~ spl24_194), inference(avatar_component_clause, [], [f1646])).
fof(f8530, plain, (~ spl24_119 | ~ spl24_104), inference(avatar_split_clause, [], [f8529, f1120, f1183])).
fof(f1120, plain, (spl24_104 <=> (e3 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_104])])).
fof(f8529, plain, (~ (e3 = op(e0, e1)) | ~ spl24_104), inference(forward_demodulation, [], [f438, f1122])).
fof(f1122, plain, ((e3 = op(e0, e4)) | ~ spl24_104), inference(avatar_component_clause, [], [f1120])).
fof(f8528, plain, (~ spl24_114 | ~ spl24_104), inference(avatar_split_clause, [], [f8527, f1120, f1162])).
fof(f8527, plain, (~ (e3 = op(e0, e2)) | ~ spl24_104), inference(forward_demodulation, [], [f440, f1122])).
fof(f8526, plain, (~ spl24_112 | ~ spl24_122), inference(avatar_split_clause, [], [f8525, f1196, f1154])).
fof(f1154, plain, (spl24_112 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_112])])).
fof(f8525, plain, (~ (e1 = op(e0, e2)) | ~ spl24_122), inference(forward_demodulation, [], [f433, f1198])).
fof(f8519, plain, (~ spl24_88 | ~ spl24_78), inference(avatar_split_clause, [], [f8518, f1011, f1053])).
fof(f1053, plain, (spl24_88 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_88])])).
fof(f8518, plain, (~ (e2 = op(e1, e2)) | ~ spl24_78), inference(forward_demodulation, [], [f450, f1013])).
fof(f1013, plain, ((e2 = op(e1, e4)) | ~ spl24_78), inference(avatar_component_clause, [], [f1011])).
fof(f8516, plain, (~ spl24_72 | ~ spl24_122), inference(avatar_split_clause, [], [f8515, f1196, f986])).
fof(f986, plain, (spl24_72 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_72])])).
fof(f8515, plain, (~ (e1 = op(e2, e0)) | ~ spl24_122), inference(forward_demodulation, [], [f383, f1198])).
fof(f8508, plain, (~ spl24_23 | ~ spl24_1 | spl24_162), inference(avatar_split_clause, [], [f8499, f1485, f688, f780])).
fof(f1485, plain, (spl24_162 <=> (e2 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_162])])).
fof(f8499, plain, (~ (e2 = op(e4, e0)) | (~ spl24_1 | spl24_162)), inference(backward_demodulation, [], [f1487, f690])).
fof(f1487, plain, (~ (e2 = op(e4, op(e4, e4))) | spl24_162), inference(avatar_component_clause, [], [f1485])).
fof(f8489, plain, (~ spl24_4 | ~ spl24_24), inference(avatar_split_clause, [], [f8487, f784, f700])).
fof(f700, plain, (spl24_4 <=> (e3 = op(e4, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_4])])).
fof(f8487, plain, (~ (e3 = op(e4, e4)) | ~ spl24_24), inference(backward_demodulation, [], [f475, f786])).
fof(f786, plain, ((e3 = op(e4, e0)) | ~ spl24_24), inference(avatar_component_clause, [], [f784])).
fof(f475, plain, ~ (op(e4, e0) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f8479, plain, (~ spl24_27 | ~ spl24_35 | spl24_176), inference(avatar_split_clause, [], [f8473, f1549, f830, f797])).
fof(f1549, plain, (spl24_176 <=> (e1 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_176])])).
fof(f8473, plain, (~ (e1 = op(e3, e4)) | (~ spl24_35 | spl24_176)), inference(backward_demodulation, [], [f1551, f832])).
fof(f1551, plain, (~ (e1 = op(e3, op(e3, e3))) | spl24_176), inference(avatar_component_clause, [], [f1549])).
fof(f8478, plain, (~ spl24_28 | ~ spl24_35 | spl24_167), inference(avatar_split_clause, [], [f8472, f1507, f830, f801])).
fof(f801, plain, (spl24_28 <=> (e2 = op(e3, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_28])])).
fof(f8472, plain, (~ (e2 = op(e3, e4)) | (~ spl24_35 | spl24_167)), inference(backward_demodulation, [], [f1509, f832])).
fof(f8477, plain, (~ spl24_6 | ~ spl24_35 | spl24_158), inference(avatar_split_clause, [], [f8470, f1468, f830, f709])).
fof(f709, plain, (spl24_6 <=> (e0 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_6])])).
fof(f1468, plain, (spl24_158 <=> (e0 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_158])])).
fof(f8470, plain, (~ (e0 = op(e4, e3)) | (~ spl24_35 | spl24_158)), inference(backward_demodulation, [], [f1470, f832])).
fof(f1470, plain, (~ (e0 = op(op(e3, e3), e3)) | spl24_158), inference(avatar_component_clause, [], [f1468])).
fof(f8467, plain, (~ spl24_28 | ~ spl24_48), inference(avatar_split_clause, [], [f8464, f885, f801])).
fof(f885, plain, (spl24_48 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_48])])).
fof(f8464, plain, (~ (e2 = op(e3, e4)) | ~ spl24_48), inference(backward_demodulation, [], [f465, f887])).
fof(f887, plain, ((e2 = op(e3, e0)) | ~ spl24_48), inference(avatar_component_clause, [], [f885])).
fof(f465, plain, ~ (op(e3, e0) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f8458, plain, (~ spl24_28 | ~ spl24_78), inference(avatar_split_clause, [], [f8451, f1011, f801])).
fof(f8451, plain, (~ (e2 = op(e3, e4)) | ~ spl24_78), inference(backward_demodulation, [], [f427, f1013])).
fof(f427, plain, ~ (op(e1, e4) = op(e3, e4)), inference(cnf_transformation, [], [f4])).
fof(f8450, plain, (~ spl24_76 | ~ spl24_81), inference(avatar_split_clause, [], [f8444, f1024, f1003])).
fof(f8444, plain, (~ (e0 = op(e1, e4)) | ~ spl24_81), inference(backward_demodulation, [], [f451, f1026])).
fof(f8449, plain, (~ spl24_6 | ~ spl24_81), inference(avatar_split_clause, [], [f8443, f1024, f709])).
fof(f8443, plain, (~ (e0 = op(e4, e3)) | ~ spl24_81), inference(backward_demodulation, [], [f418, f1026])).
fof(f8440, plain, (~ spl24_78 | ~ spl24_95 | spl24_182), inference(avatar_split_clause, [], [f8435, f1580, f1082, f1011])).
fof(f1580, plain, (spl24_182 <=> (e2 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl24_182])])).
fof(f8435, plain, (~ (e2 = op(e1, e4)) | (~ spl24_95 | spl24_182)), inference(backward_demodulation, [], [f1582, f1084])).
fof(f1582, plain, (~ (e2 = op(e1, op(e1, e1))) | spl24_182), inference(avatar_component_clause, [], [f1580])).
fof(f8426, plain, (~ spl24_94 | ~ spl24_99), inference(avatar_split_clause, [], [f8421, f1099, f1078])).
fof(f8421, plain, (~ (e3 = op(e1, e1)) | ~ spl24_99), inference(backward_demodulation, [], [f442, f1101])).
fof(f8425, plain, (~ spl24_24 | ~ spl24_99), inference(avatar_split_clause, [], [f8420, f1099, f784])).
fof(f8420, plain, (~ (e3 = op(e4, e0)) | ~ spl24_99), inference(backward_demodulation, [], [f388, f1101])).
fof(f8418, plain, (~ spl24_4 | ~ spl24_104), inference(avatar_split_clause, [], [f8416, f1120, f700])).
fof(f8416, plain, (~ (e3 = op(e4, e4)) | ~ spl24_104), inference(backward_demodulation, [], [f425, f1122])).
fof(f425, plain, ~ (op(e0, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f8407, plain, (~ spl24_122 | ~ spl24_125), inference(avatar_contradiction_clause, [], [f8406])).
fof(f8406, plain, ($false | (~ spl24_122 | ~ spl24_125)), inference(subsumption_resolution, [], [f8405, f488])).
fof(f488, plain, ~ (e1 = e4), inference(cnf_transformation, [], [f5])).
fof(f8405, plain, ((e1 = e4) | (~ spl24_122 | ~ spl24_125)), inference(backward_demodulation, [], [f1210, f1198])).
fof(f1210, plain, ((op(e0, e0) = e4) | ~ spl24_125), inference(avatar_component_clause, [], [f1208])).
fof(f1208, plain, (spl24_125 <=> (op(e0, e0) = e4)), introduced(avatar_definition, [new_symbols(naming, [spl24_125])])).
fof(f8401, plain, (~ spl24_102 | ~ spl24_125 | spl24_206), inference(avatar_split_clause, [], [f8395, f1724, f1208, f1112])).
fof(f1724, plain, (spl24_206 <=> (e1 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_206])])).
fof(f8395, plain, (~ (e1 = op(e0, e4)) | (~ spl24_125 | spl24_206)), inference(backward_demodulation, [], [f1726, f1210])).
fof(f1726, plain, (~ (e1 = op(e0, op(e0, e0))) | spl24_206), inference(avatar_component_clause, [], [f1724])).
fof(f8398, plain, (~ spl24_120 | ~ spl24_125), inference(avatar_split_clause, [], [f8389, f1208, f1187])).
fof(f8389, plain, (~ (e4 = op(e0, e1)) | ~ spl24_125), inference(backward_demodulation, [], [f432, f1210])).
fof(f432, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f4])).
fof(f8374, plain, (~ spl24_50 | spl24_193 | ~ spl24_194), inference(avatar_split_clause, [], [f8371, f1646, f1642, f893])).
fof(f893, plain, (spl24_50 <=> (e4 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_50])])).
fof(f1642, plain, (spl24_193 <=> (e4 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_193])])).
fof(f8371, plain, (~ (e4 = op(e3, e0)) | (spl24_193 | ~ spl24_194)), inference(backward_demodulation, [], [f1644, f1647])).
fof(f1644, plain, (~ (e4 = op(op(e0, op(e0, e0)), e0)) | spl24_193), inference(avatar_component_clause, [], [f1642])).
fof(f8368, plain, (~ spl24_210 | ~ spl24_214), inference(avatar_contradiction_clause, [], [f8367])).
fof(f8367, plain, ($false | (~ spl24_210 | ~ spl24_214)), inference(subsumption_resolution, [], [f8366, f491])).
fof(f491, plain, ~ (e3 = e4), inference(cnf_transformation, [], [f5])).
fof(f8366, plain, ((e3 = e4) | (~ spl24_210 | ~ spl24_214)), inference(backward_demodulation, [], [f1797, f1757])).
fof(f1757, plain, ((e3 = op(op(e0, e0), e0)) | ~ spl24_210), inference(avatar_component_clause, [], [f1756])).
fof(f1797, plain, ((e4 = op(op(e0, e0), e0)) | ~ spl24_214), inference(avatar_component_clause, [], [f1796])).
fof(f8355, plain, (~ spl24_103 | ~ spl24_108), inference(avatar_split_clause, [], [f8354, f1137, f1116])).
fof(f1116, plain, (spl24_103 <=> (e2 = op(e0, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_103])])).
fof(f8354, plain, (~ (e2 = op(e0, e4)) | ~ spl24_108), inference(forward_demodulation, [], [f441, f1139])).
fof(f1139, plain, ((e2 = op(e0, e3)) | ~ spl24_108), inference(avatar_component_clause, [], [f1137])).
fof(f8344, plain, (~ spl24_77 | ~ spl24_87), inference(avatar_split_clause, [], [f8343, f1049, f1007])).
fof(f1049, plain, (spl24_87 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_87])])).
fof(f8343, plain, (~ (e1 = op(e1, e4)) | ~ spl24_87), inference(forward_demodulation, [], [f450, f1051])).
fof(f1051, plain, ((e1 = op(e1, e2)) | ~ spl24_87), inference(avatar_component_clause, [], [f1049])).
fof(f8328, plain, (~ spl24_10 | ~ spl24_15), inference(avatar_split_clause, [], [f8326, f746, f725])).
fof(f8326, plain, (~ (e4 = op(e4, e3)) | ~ spl24_15), inference(backward_demodulation, [], [f479, f748])).
fof(f748, plain, ((e4 = op(e4, e2)) | ~ spl24_15), inference(avatar_component_clause, [], [f746])).
fof(f479, plain, ~ (op(e4, e2) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f8325, plain, (~ spl24_3 | ~ spl24_18), inference(avatar_split_clause, [], [f8322, f759, f696])).
fof(f8322, plain, (~ (e2 = op(e4, e4)) | ~ spl24_18), inference(backward_demodulation, [], [f478, f761])).
fof(f761, plain, ((e2 = op(e4, e1)) | ~ spl24_18), inference(avatar_component_clause, [], [f759])).
fof(f8319, plain, (~ spl24_12 | ~ spl24_22), inference(avatar_split_clause, [], [f8313, f776, f734])).
fof(f8313, plain, (~ (e1 = op(e4, e2)) | ~ spl24_22), inference(backward_demodulation, [], [f473, f778])).
fof(f473, plain, ~ (op(e4, e0) = op(e4, e2)), inference(cnf_transformation, [], [f4])).
fof(f8311, plain, (~ spl24_3 | ~ spl24_28), inference(avatar_split_clause, [], [f8308, f801, f696])).
fof(f8308, plain, (~ (e2 = op(e4, e4)) | ~ spl24_28), inference(backward_demodulation, [], [f431, f803])).
fof(f803, plain, ((e2 = op(e3, e4)) | ~ spl24_28), inference(avatar_component_clause, [], [f801])).
fof(f8307, plain, (~ spl24_32 | ~ spl24_85 | spl24_211), inference(avatar_contradiction_clause, [], [f8306])).
fof(f8306, plain, ($false | (~ spl24_32 | ~ spl24_85 | spl24_211)), inference(subsumption_resolution, [], [f8305, f1042])).
fof(f8305, plain, (~ (e4 = op(e1, e3)) | (~ spl24_32 | spl24_211)), inference(backward_demodulation, [], [f1777, f820])).
fof(f8297, plain, (~ spl24_34 | ~ spl24_39), inference(avatar_split_clause, [], [f8294, f847, f826])).
fof(f8294, plain, (~ (e3 = op(e3, e3)) | ~ spl24_39), inference(backward_demodulation, [], [f469, f849])).
fof(f849, plain, ((e3 = op(e3, e2)) | ~ spl24_39), inference(avatar_component_clause, [], [f847])).
fof(f8292, plain, (~ spl24_16 | ~ spl24_41), inference(avatar_split_clause, [], [f8284, f856, f751])).
fof(f8284, plain, (~ (e0 = op(e4, e1)) | ~ spl24_41), inference(backward_demodulation, [], [f401, f858])).
fof(f858, plain, ((e0 = op(e3, e1)) | ~ spl24_41), inference(avatar_component_clause, [], [f856])).
fof(f8283, plain, (~ spl24_30 | ~ spl24_50), inference(avatar_split_clause, [], [f8282, f893, f809])).
fof(f8282, plain, (~ (e4 = op(e3, e4)) | ~ spl24_50), inference(backward_demodulation, [], [f465, f895])).
fof(f895, plain, ((e4 = op(e3, e0)) | ~ spl24_50), inference(avatar_component_clause, [], [f893])).
fof(f8277, plain, (~ spl24_30 | ~ spl24_55), inference(avatar_split_clause, [], [f8275, f914, f809])).
fof(f8275, plain, (~ (e4 = op(e3, e4)) | ~ spl24_55), inference(backward_demodulation, [], [f429, f916])).
fof(f916, plain, ((e4 = op(e2, e4)) | ~ spl24_55), inference(avatar_component_clause, [], [f914])).
fof(f8257, plain, (~ spl24_58 | ~ spl24_63), inference(avatar_split_clause, [], [f8243, f948, f927])).
fof(f948, plain, (spl24_63 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_63])])).
fof(f8243, plain, (~ (e2 = op(e2, e3)) | ~ spl24_63), inference(backward_demodulation, [], [f459, f950])).
fof(f950, plain, ((e2 = op(e2, e2)) | ~ spl24_63), inference(avatar_component_clause, [], [f948])).
fof(f459, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f8256, plain, (~ spl24_38 | ~ spl24_63), inference(avatar_split_clause, [], [f8241, f948, f843])).
fof(f8241, plain, (~ (e2 = op(e3, e2)) | ~ spl24_63), inference(backward_demodulation, [], [f409, f950])).
fof(f409, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f8240, plain, (~ spl24_42 | ~ spl24_67), inference(avatar_split_clause, [], [f8232, f965, f860])).
fof(f965, plain, (spl24_67 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_67])])).
fof(f8232, plain, (~ (e1 = op(e3, e1)) | ~ spl24_67), inference(backward_demodulation, [], [f399, f967])).
fof(f967, plain, ((e1 = op(e2, e1)) | ~ spl24_67), inference(avatar_component_clause, [], [f965])).
fof(f399, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f8231, plain, (~ spl24_61 | ~ spl24_71), inference(avatar_split_clause, [], [f8223, f982, f940])).
fof(f982, plain, (spl24_71 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_71])])).
fof(f8223, plain, (~ (e0 = op(e2, e2)) | ~ spl24_71), inference(backward_demodulation, [], [f453, f984])).
fof(f984, plain, ((e0 = op(e2, e0)) | ~ spl24_71), inference(avatar_component_clause, [], [f982])).
fof(f453, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f8230, plain, (~ spl24_46 | ~ spl24_71), inference(avatar_split_clause, [], [f8220, f982, f877])).
fof(f8220, plain, (~ (e0 = op(e3, e0)) | ~ spl24_71), inference(backward_demodulation, [], [f389, f984])).
fof(f389, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f8215, plain, (~ spl24_82 | ~ spl24_87), inference(avatar_split_clause, [], [f8210, f1049, f1028])).
fof(f8210, plain, (~ (e1 = op(e1, e3)) | ~ spl24_87), inference(backward_demodulation, [], [f449, f1051])).
fof(f449, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f4])).
fof(f8192, plain, (~ spl24_93 | ~ spl24_98), inference(avatar_split_clause, [], [f8187, f1095, f1074])).
fof(f1095, plain, (spl24_98 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_98])])).
fof(f8187, plain, (~ (e2 = op(e1, e1)) | ~ spl24_98), inference(backward_demodulation, [], [f442, f1097])).
fof(f1097, plain, ((e2 = op(e1, e0)) | ~ spl24_98), inference(avatar_component_clause, [], [f1095])).
fof(f8183, plain, (~ spl24_33 | ~ spl24_108), inference(avatar_split_clause, [], [f8178, f1137, f822])).
fof(f8178, plain, (~ (e2 = op(e3, e3)) | ~ spl24_108), inference(backward_demodulation, [], [f414, f1139])).
fof(f8182, plain, (~ spl24_58 | ~ spl24_108), inference(avatar_split_clause, [], [f8177, f1137, f927])).
fof(f8177, plain, (~ (e2 = op(e2, e3)) | ~ spl24_108), inference(backward_demodulation, [], [f413, f1139])).
fof(f413, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f4])).
fof(f8175, plain, (~ spl24_106 | ~ spl24_111), inference(avatar_split_clause, [], [f8170, f1150, f1129])).
fof(f1150, plain, (spl24_111 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_111])])).
fof(f8170, plain, (~ (e0 = op(e0, e3)) | ~ spl24_111), inference(backward_demodulation, [], [f439, f1152])).
fof(f1152, plain, ((e0 = op(e0, e2)) | ~ spl24_111), inference(avatar_component_clause, [], [f1150])).
fof(f439, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f4])).
fof(f8174, plain, (~ spl24_61 | ~ spl24_111), inference(avatar_split_clause, [], [f8167, f1150, f940])).
fof(f8167, plain, (~ (e0 = op(e2, e2)) | ~ spl24_111), inference(backward_demodulation, [], [f403, f1152])).
fof(f403, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f4])).
fof(f8164, plain, (~ spl24_70 | ~ spl24_120), inference(avatar_split_clause, [], [f8159, f1187, f977])).
fof(f8159, plain, (~ (e4 = op(e2, e1)) | ~ spl24_120), inference(backward_demodulation, [], [f393, f1189])).
fof(f393, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f8156, plain, (spl24_108 | ~ spl24_124 | ~ spl24_198), inference(avatar_split_clause, [], [f8149, f1668, f1204, f1137])).
fof(f1204, plain, (spl24_124 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl24_124])])).
fof(f1668, plain, (spl24_198 <=> (e2 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl24_198])])).
fof(f8149, plain, ((e2 = op(e0, e3)) | (~ spl24_124 | ~ spl24_198)), inference(backward_demodulation, [], [f1669, f1206])).
fof(f1206, plain, ((op(e0, e0) = e3) | ~ spl24_124), inference(avatar_component_clause, [], [f1204])).
fof(f1669, plain, ((e2 = op(e0, op(e0, e0))) | ~ spl24_198), inference(avatar_component_clause, [], [f1668])).
fof(f8155, plain, (~ spl24_119 | ~ spl24_124), inference(avatar_split_clause, [], [f8144, f1204, f1183])).
fof(f8144, plain, (~ (e3 = op(e0, e1)) | ~ spl24_124), inference(backward_demodulation, [], [f432, f1206])).
fof(f8154, plain, (~ spl24_24 | ~ spl24_124), inference(avatar_split_clause, [], [f8143, f1204, f784])).
fof(f8143, plain, (~ (e3 = op(e4, e0)) | ~ spl24_124), inference(backward_demodulation, [], [f385, f1206])).
fof(f385, plain, ~ (op(e0, e0) = op(e4, e0)), inference(cnf_transformation, [], [f4])).
fof(f8139, plain, (spl24_15 | ~ spl24_128), inference(avatar_split_clause, [], [f8127, f1221, f746])).
fof(f1221, plain, (spl24_128 <=> (e2 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl24_128])])).
fof(f8127, plain, ((e4 = op(e4, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f330, f1223])).
fof(f1223, plain, ((e2 = unit) | ~ spl24_128), inference(avatar_component_clause, [], [f1221])).
fof(f8138, plain, (spl24_55 | ~ spl24_128), inference(avatar_split_clause, [], [f8126, f1221, f914])).
fof(f8126, plain, ((e4 = op(e2, e4)) | ~ spl24_128), inference(backward_demodulation, [], [f329, f1223])).
fof(f8137, plain, (spl24_39 | ~ spl24_128), inference(avatar_split_clause, [], [f8125, f1221, f847])).
fof(f8125, plain, ((e3 = op(e3, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f328, f1223])).
fof(f8136, plain, (spl24_59 | ~ spl24_128), inference(avatar_split_clause, [], [f8124, f1221, f931])).
fof(f8124, plain, ((e3 = op(e2, e3)) | ~ spl24_128), inference(backward_demodulation, [], [f327, f1223])).
fof(f8135, plain, (spl24_63 | ~ spl24_128), inference(avatar_split_clause, [], [f8123, f1221, f948])).
fof(f8123, plain, ((e2 = op(e2, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f326, f1223])).
fof(f8133, plain, (spl24_87 | ~ spl24_128), inference(avatar_split_clause, [], [f8121, f1221, f1049])).
fof(f8121, plain, ((e1 = op(e1, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f324, f1223])).
fof(f8132, plain, (spl24_67 | ~ spl24_128), inference(avatar_split_clause, [], [f8120, f1221, f965])).
fof(f8120, plain, ((e1 = op(e2, e1)) | ~ spl24_128), inference(backward_demodulation, [], [f323, f1223])).
fof(f8131, plain, (spl24_111 | ~ spl24_128), inference(avatar_split_clause, [], [f8119, f1221, f1150])).
fof(f8119, plain, ((e0 = op(e0, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f322, f1223])).
fof(f8130, plain, (spl24_71 | ~ spl24_128), inference(avatar_split_clause, [], [f8118, f1221, f982])).
fof(f8118, plain, ((e0 = op(e2, e0)) | ~ spl24_128), inference(backward_demodulation, [], [f321, f1223])).
fof(f8111, plain, (~ spl24_18 | ~ spl24_175 | spl24_180), inference(avatar_split_clause, [], [f8110, f1570, f1544, f759])).
fof(f8110, plain, (~ (e2 = op(e4, e1)) | (~ spl24_175 | spl24_180)), inference(backward_demodulation, [], [f1572, f1545])).
fof(f8095, plain, (~ spl24_72 | ~ spl24_198 | spl24_204), inference(avatar_split_clause, [], [f8092, f1714, f1668, f986])).
fof(f1714, plain, (spl24_204 <=> (e1 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_204])])).
fof(f8092, plain, (~ (e1 = op(e2, e0)) | (~ spl24_198 | spl24_204)), inference(backward_demodulation, [], [f1716, f1669])).
fof(f1716, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | spl24_204), inference(avatar_component_clause, [], [f1714])).
fof(f8000, plain, (~ spl24_57 | ~ spl24_72), inference(avatar_split_clause, [], [f7991, f986, f923])).
fof(f923, plain, (spl24_57 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_57])])).
fof(f7991, plain, (~ (e1 = op(e2, e3)) | ~ spl24_72), inference(backward_demodulation, [], [f454, f988])).
fof(f988, plain, ((e1 = op(e2, e0)) | ~ spl24_72), inference(avatar_component_clause, [], [f986])).
fof(f7999, plain, (~ spl24_47 | ~ spl24_72), inference(avatar_split_clause, [], [f7989, f986, f881])).
fof(f7989, plain, (~ (e1 = op(e3, e0)) | ~ spl24_72), inference(backward_demodulation, [], [f389, f988])).
fof(f7980, plain, (~ spl24_57 | ~ spl24_82), inference(avatar_split_clause, [], [f7974, f1028, f923])).
fof(f7974, plain, (~ (e1 = op(e2, e3)) | ~ spl24_82), inference(backward_demodulation, [], [f416, f1030])).
fof(f7972, plain, (~ spl24_69 | ~ spl24_93 | spl24_209), inference(avatar_split_clause, [], [f7968, f1746, f1074, f973])).
fof(f973, plain, (spl24_69 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_69])])).
fof(f1746, plain, (spl24_209 <=> (e3 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_209])])).
fof(f7968, plain, (~ (e3 = op(e2, e1)) | (~ spl24_93 | spl24_209)), inference(backward_demodulation, [], [f1748, f1076])).
fof(f1748, plain, (~ (e3 = op(op(e1, e1), e1)) | spl24_209), inference(avatar_component_clause, [], [f1746])).
fof(f7956, plain, (~ spl24_75 | ~ spl24_100), inference(avatar_split_clause, [], [f7951, f1103, f998])).
fof(f7951, plain, (~ (e4 = op(e2, e0)) | ~ spl24_100), inference(backward_demodulation, [], [f386, f1105])).
fof(f7930, plain, (~ spl24_110 | ~ spl24_115), inference(avatar_split_clause, [], [f7926, f1166, f1145])).
fof(f1145, plain, (spl24_110 <=> (e4 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_110])])).
fof(f7926, plain, (~ (e4 = op(e0, e3)) | ~ spl24_115), inference(backward_demodulation, [], [f439, f1168])).
fof(f7921, plain, (~ spl24_33 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f7920])).
fof(f7920, plain, ($false | (~ spl24_33 | ~ spl24_129)), inference(subsumption_resolution, [], [f7919, f489])).
fof(f489, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f5])).
fof(f7919, plain, ((e2 = e3) | (~ spl24_33 | ~ spl24_129)), inference(forward_demodulation, [], [f7906, f824])).
fof(f7906, plain, ((e3 = op(e3, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f328, f1227])).
fof(f7893, plain, (~ spl24_75 | ~ spl24_61 | spl24_165), inference(avatar_split_clause, [], [f7892, f1498, f940, f998])).
fof(f7892, plain, (~ (e4 = op(e2, e0)) | (~ spl24_61 | spl24_165)), inference(forward_demodulation, [], [f1500, f942])).
fof(f942, plain, ((e0 = op(e2, e2)) | ~ spl24_61), inference(avatar_component_clause, [], [f940])).
fof(f1500, plain, (~ (e4 = op(e2, op(e2, e2))) | spl24_165), inference(avatar_component_clause, [], [f1498])).
fof(f7887, plain, (~ spl24_112 | ~ spl24_61 | spl24_185), inference(avatar_split_clause, [], [f7886, f1602, f940, f1154])).
fof(f1602, plain, (spl24_185 <=> (e1 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_185])])).
fof(f7886, plain, (~ (e1 = op(e0, e2)) | (~ spl24_61 | spl24_185)), inference(forward_demodulation, [], [f1604, f942])).
fof(f1604, plain, (~ (e1 = op(op(e2, e2), e2)) | spl24_185), inference(avatar_component_clause, [], [f1602])).
fof(f7883, plain, (~ spl24_112 | ~ spl24_123 | spl24_206), inference(avatar_split_clause, [], [f7882, f1724, f1200, f1154])).
fof(f7882, plain, (~ (e1 = op(e0, e2)) | (~ spl24_123 | spl24_206)), inference(forward_demodulation, [], [f1726, f1202])).
fof(f7881, plain, (~ spl24_115 | ~ spl24_61 | spl24_212), inference(avatar_split_clause, [], [f7880, f1780, f940, f1166])).
fof(f7880, plain, (~ (e4 = op(e0, e2)) | (~ spl24_61 | spl24_212)), inference(forward_demodulation, [], [f1782, f942])).
fof(f7879, plain, (~ spl24_75 | ~ spl24_123 | spl24_214), inference(avatar_split_clause, [], [f7878, f1796, f1200, f998])).
fof(f7878, plain, (~ (e4 = op(e2, e0)) | (~ spl24_123 | spl24_214)), inference(forward_demodulation, [], [f1798, f1202])).
fof(f7860, plain, (~ spl24_69 | ~ spl24_119), inference(avatar_split_clause, [], [f7859, f1183, f973])).
fof(f7859, plain, (~ (e3 = op(e2, e1)) | ~ spl24_119), inference(forward_demodulation, [], [f393, f1185])).
fof(f7856, plain, (~ spl24_119 | spl24_174 | ~ spl24_205), inference(avatar_contradiction_clause, [], [f7855])).
fof(f7855, plain, ($false | (~ spl24_119 | spl24_174 | ~ spl24_205)), inference(subsumption_resolution, [], [f7684, f1185])).
fof(f7684, plain, (~ (e3 = op(e0, e1)) | (spl24_174 | ~ spl24_205)), inference(backward_demodulation, [], [f1542, f1720])).
fof(f1720, plain, ((e0 = op(e1, op(e1, e1))) | ~ spl24_205), inference(avatar_component_clause, [], [f1719])).
fof(f7854, plain, (~ spl24_118 | spl24_180 | ~ spl24_205), inference(avatar_split_clause, [], [f7689, f1719, f1570, f1179])).
fof(f7689, plain, (~ (e2 = op(e0, e1)) | (spl24_180 | ~ spl24_205)), inference(forward_demodulation, [], [f1572, f1720])).
fof(f7845, plain, (~ spl24_103 | ~ spl24_123), inference(avatar_split_clause, [], [f7844, f1200, f1116])).
fof(f7844, plain, (~ (e2 = op(e0, e4)) | ~ spl24_123), inference(forward_demodulation, [], [f435, f1202])).
fof(f435, plain, ~ (op(e0, e0) = op(e0, e4)), inference(cnf_transformation, [], [f4])).
fof(f7841, plain, (~ spl24_98 | ~ spl24_123), inference(avatar_split_clause, [], [f7840, f1200, f1095])).
fof(f7840, plain, (~ (e2 = op(e1, e0)) | ~ spl24_123), inference(forward_demodulation, [], [f382, f1202])).
fof(f382, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f4])).
fof(f7837, plain, (~ spl24_84 | ~ spl24_89), inference(avatar_split_clause, [], [f7836, f1057, f1036])).
fof(f1036, plain, (spl24_84 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_84])])).
fof(f7836, plain, (~ (e3 = op(e1, e3)) | ~ spl24_89), inference(forward_demodulation, [], [f449, f1059])).
fof(f7822, plain, (~ spl24_48 | ~ spl24_123), inference(avatar_split_clause, [], [f7821, f1200, f885])).
fof(f7821, plain, (~ (e2 = op(e3, e0)) | ~ spl24_123), inference(forward_demodulation, [], [f384, f1202])).
fof(f384, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f7817, plain, (~ spl24_9 | ~ spl24_24), inference(avatar_split_clause, [], [f7816, f784, f721])).
fof(f7816, plain, (~ (e3 = op(e4, e3)) | ~ spl24_24), inference(forward_demodulation, [], [f474, f786])).
fof(f7810, plain, (~ spl24_22 | ~ spl24_191 | spl24_204), inference(avatar_split_clause, [], [f7702, f1714, f1632, f776])).
fof(f7702, plain, (~ (e1 = op(e4, e0)) | (~ spl24_191 | spl24_204)), inference(forward_demodulation, [], [f1716, f1633])).
fof(f7806, plain, (~ spl24_21 | ~ spl24_24), inference(avatar_contradiction_clause, [], [f7805])).
fof(f7805, plain, ($false | (~ spl24_21 | ~ spl24_24)), inference(subsumption_resolution, [], [f7804, f484])).
fof(f7804, plain, ((e0 = e3) | (~ spl24_21 | ~ spl24_24)), inference(forward_demodulation, [], [f786, f774])).
fof(f7803, plain, (~ spl24_9 | ~ spl24_33 | ~ spl24_40 | ~ spl24_166), inference(avatar_contradiction_clause, [], [f7802])).
fof(f7802, plain, ($false | (~ spl24_9 | ~ spl24_33 | ~ spl24_40 | ~ spl24_166)), inference(subsumption_resolution, [], [f7801, f491])).
fof(f7801, plain, ((e3 = e4) | (~ spl24_9 | ~ spl24_33 | ~ spl24_40 | ~ spl24_166)), inference(forward_demodulation, [], [f7797, f723])).
fof(f723, plain, ((e3 = op(e4, e3)) | ~ spl24_9), inference(avatar_component_clause, [], [f721])).
fof(f7797, plain, ((e4 = op(e4, e3)) | (~ spl24_33 | ~ spl24_40 | ~ spl24_166)), inference(backward_demodulation, [], [f7715, f853])).
fof(f7715, plain, ((e4 = op(op(e3, e2), e3)) | (~ spl24_33 | ~ spl24_166)), inference(forward_demodulation, [], [f1504, f824])).
fof(f1504, plain, ((e4 = op(op(e3, op(e3, e3)), e3)) | ~ spl24_166), inference(avatar_component_clause, [], [f1503])).
fof(f1503, plain, (spl24_166 <=> (e4 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_166])])).
fof(f7796, plain, (~ spl24_36 | ~ spl24_41), inference(avatar_split_clause, [], [f7791, f856, f835])).
fof(f7791, plain, (~ (e0 = op(e3, e2)) | ~ spl24_41), inference(backward_demodulation, [], [f466, f858])).
fof(f7788, plain, (~ spl24_36 | ~ spl24_61), inference(avatar_split_clause, [], [f7778, f940, f835])).
fof(f7778, plain, (~ (e0 = op(e3, e2)) | ~ spl24_61), inference(backward_demodulation, [], [f409, f942])).
fof(f7777, plain, (~ spl24_64 | ~ spl24_69), inference(avatar_split_clause, [], [f7774, f973, f952])).
fof(f952, plain, (spl24_64 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_64])])).
fof(f7774, plain, (~ (e3 = op(e2, e2)) | ~ spl24_69), inference(backward_demodulation, [], [f456, f975])).
fof(f975, plain, ((e3 = op(e2, e1)) | ~ spl24_69), inference(avatar_component_clause, [], [f973])).
fof(f7773, plain, (~ spl24_64 | ~ spl24_89), inference(avatar_split_clause, [], [f7770, f1057, f952])).
fof(f7770, plain, (~ (e3 = op(e2, e2)) | ~ spl24_89), inference(backward_demodulation, [], [f406, f1059])).
fof(f7761, plain, (~ spl24_45 | ~ spl24_95), inference(avatar_split_clause, [], [f7754, f1082, f872])).
fof(f7754, plain, (~ (e4 = op(e3, e1)) | ~ spl24_95), inference(backward_demodulation, [], [f397, f1084])).
fof(f7752, plain, (~ spl24_118 | ~ spl24_119), inference(avatar_contradiction_clause, [], [f7751])).
fof(f7751, plain, ($false | (~ spl24_118 | ~ spl24_119)), inference(subsumption_resolution, [], [f7750, f489])).
fof(f7750, plain, ((e2 = e3) | (~ spl24_118 | ~ spl24_119)), inference(forward_demodulation, [], [f1185, f1181])).
fof(f7727, plain, (~ spl24_36 | ~ spl24_33 | spl24_192), inference(avatar_split_clause, [], [f7726, f1637, f822, f835])).
fof(f1637, plain, (spl24_192 <=> (e0 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl24_192])])).
fof(f7726, plain, (~ (e0 = op(e3, e2)) | (~ spl24_33 | spl24_192)), inference(forward_demodulation, [], [f1639, f824])).
fof(f1639, plain, (~ (e0 = op(e3, op(e3, e3))) | spl24_192), inference(avatar_component_clause, [], [f1637])).
fof(f7709, plain, (~ spl24_123 | ~ spl24_118), inference(avatar_split_clause, [], [f7708, f1179, f1200])).
fof(f7708, plain, (~ (op(e0, e0) = e2) | ~ spl24_118), inference(forward_demodulation, [], [f432, f1181])).
fof(f7698, plain, (~ spl24_92 | ~ spl24_17), inference(avatar_split_clause, [], [f7697, f755, f1070])).
fof(f1070, plain, (spl24_92 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_92])])).
fof(f7697, plain, (~ (e1 = op(e1, e1)) | ~ spl24_17), inference(forward_demodulation, [], [f398, f757])).
fof(f7695, plain, (~ spl24_92 | ~ spl24_77), inference(avatar_split_clause, [], [f7694, f1007, f1070])).
fof(f7694, plain, (~ (e1 = op(e1, e1)) | ~ spl24_77), inference(forward_demodulation, [], [f448, f1009])).
fof(f1009, plain, ((e1 = op(e1, e4)) | ~ spl24_77), inference(avatar_component_clause, [], [f1007])).
fof(f7652, plain, (~ spl24_4 | ~ spl24_9), inference(avatar_split_clause, [], [f7650, f721, f700])).
fof(f7650, plain, (~ (e3 = op(e4, e4)) | ~ spl24_9), inference(backward_demodulation, [], [f481, f723])).
fof(f7629, plain, (~ spl24_4 | ~ spl24_29), inference(avatar_split_clause, [], [f7627, f805, f700])).
fof(f7627, plain, (~ (e3 = op(e4, e4)) | ~ spl24_29), inference(backward_demodulation, [], [f431, f807])).
fof(f807, plain, ((e3 = op(e3, e4)) | ~ spl24_29), inference(avatar_component_clause, [], [f805])).
fof(f7616, plain, (~ spl24_31 | ~ spl24_36), inference(avatar_split_clause, [], [f7610, f835, f814])).
fof(f814, plain, (spl24_31 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_31])])).
fof(f7610, plain, (~ (e0 = op(e3, e3)) | ~ spl24_36), inference(backward_demodulation, [], [f469, f837])).
fof(f7608, plain, (~ spl24_27 | ~ spl24_47), inference(avatar_split_clause, [], [f7603, f881, f797])).
fof(f7603, plain, (~ (e1 = op(e3, e4)) | ~ spl24_47), inference(backward_demodulation, [], [f465, f883])).
fof(f7607, plain, (~ spl24_22 | ~ spl24_47), inference(avatar_split_clause, [], [f7600, f881, f776])).
fof(f7600, plain, (~ (e1 = op(e4, e0)) | ~ spl24_47), inference(backward_demodulation, [], [f391, f883])).
fof(f7589, plain, (spl24_57 | ~ spl24_64 | ~ spl24_181), inference(avatar_split_clause, [], [f7584, f1575, f952, f923])).
fof(f7584, plain, ((e1 = op(e2, e3)) | (~ spl24_64 | ~ spl24_181)), inference(backward_demodulation, [], [f1576, f954])).
fof(f954, plain, ((e3 = op(e2, e2)) | ~ spl24_64), inference(avatar_component_clause, [], [f952])).
fof(f7587, plain, (~ spl24_59 | ~ spl24_64), inference(avatar_split_clause, [], [f7580, f952, f931])).
fof(f7580, plain, (~ (e3 = op(e2, e3)) | ~ spl24_64), inference(backward_demodulation, [], [f459, f954])).
fof(f7586, plain, (~ spl24_39 | ~ spl24_64), inference(avatar_split_clause, [], [f7578, f952, f847])).
fof(f7578, plain, (~ (e3 = op(e3, e2)) | ~ spl24_64), inference(backward_demodulation, [], [f409, f954])).
fof(f7568, plain, (~ spl24_55 | ~ spl24_75), inference(avatar_split_clause, [], [f7567, f998, f914])).
fof(f7567, plain, (~ (e4 = op(e2, e4)) | ~ spl24_75), inference(backward_demodulation, [], [f455, f1000])).
fof(f7561, plain, (~ spl24_27 | ~ spl24_77), inference(avatar_split_clause, [], [f7556, f1007, f797])).
fof(f7556, plain, (~ (e1 = op(e3, e4)) | ~ spl24_77), inference(backward_demodulation, [], [f427, f1009])).
fof(f7552, plain, (~ spl24_81 | ~ spl24_94 | ~ spl24_118 | spl24_180), inference(avatar_contradiction_clause, [], [f7551])).
fof(f7551, plain, ($false | (~ spl24_81 | ~ spl24_94 | ~ spl24_118 | spl24_180)), inference(subsumption_resolution, [], [f7546, f1181])).
fof(f7546, plain, (~ (e2 = op(e0, e1)) | (~ spl24_81 | ~ spl24_94 | spl24_180)), inference(backward_demodulation, [], [f7375, f1026])).
fof(f7375, plain, (~ (e2 = op(op(e1, e3), e1)) | (~ spl24_94 | spl24_180)), inference(backward_demodulation, [], [f1572, f1080])).
fof(f7549, plain, (~ spl24_31 | ~ spl24_81), inference(avatar_split_clause, [], [f7539, f1024, f814])).
fof(f7539, plain, (~ (e0 = op(e3, e3)) | ~ spl24_81), inference(backward_demodulation, [], [f417, f1026])).
fof(f417, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f7537, plain, (~ spl24_15 | ~ spl24_90), inference(avatar_split_clause, [], [f7534, f1061, f746])).
fof(f7534, plain, (~ (e4 = op(e4, e2)) | ~ spl24_90), inference(backward_demodulation, [], [f408, f1063])).
fof(f7531, plain, (~ spl24_78 | ~ spl24_98), inference(avatar_split_clause, [], [f7527, f1095, f1011])).
fof(f7527, plain, (~ (e2 = op(e1, e4)) | ~ spl24_98), inference(backward_demodulation, [], [f445, f1097])).
fof(f445, plain, ~ (op(e1, e0) = op(e1, e4)), inference(cnf_transformation, [], [f4])).
fof(f7530, plain, (~ spl24_48 | ~ spl24_98), inference(avatar_split_clause, [], [f7523, f1095, f885])).
fof(f7523, plain, (~ (e2 = op(e3, e0)) | ~ spl24_98), inference(backward_demodulation, [], [f387, f1097])).
fof(f387, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f4])).
fof(f7511, plain, (~ spl24_87 | ~ spl24_112), inference(avatar_split_clause, [], [f7503, f1154, f1049])).
fof(f7503, plain, (~ (e1 = op(e1, e2)) | ~ spl24_112), inference(backward_demodulation, [], [f402, f1156])).
fof(f1156, plain, ((e1 = op(e0, e2)) | ~ spl24_112), inference(avatar_component_clause, [], [f1154])).
fof(f7464, plain, (~ spl24_48 | ~ spl24_124 | spl24_202), inference(avatar_split_clause, [], [f7463, f1700, f1204, f885])).
fof(f1700, plain, (spl24_202 <=> (e2 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_202])])).
fof(f7463, plain, (~ (e2 = op(e3, e0)) | (~ spl24_124 | spl24_202)), inference(forward_demodulation, [], [f1702, f1206])).
fof(f1702, plain, (~ (e2 = op(op(e0, e0), e0)) | spl24_202), inference(avatar_component_clause, [], [f1700])).
fof(f7458, plain, (~ spl24_113 | ~ spl24_118), inference(avatar_split_clause, [], [f7457, f1179, f1158])).
fof(f7457, plain, (~ (e2 = op(e0, e2)) | ~ spl24_118), inference(forward_demodulation, [], [f436, f1181])).
fof(f7452, plain, (~ spl24_69 | ~ spl24_94), inference(avatar_split_clause, [], [f7451, f1078, f973])).
fof(f7451, plain, (~ (e3 = op(e2, e1)) | ~ spl24_94), inference(forward_demodulation, [], [f396, f1080])).
fof(f7450, plain, (~ spl24_68 | ~ spl24_118), inference(avatar_split_clause, [], [f7449, f1179, f969])).
fof(f969, plain, (spl24_68 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_68])])).
fof(f7449, plain, (~ (e2 = op(e2, e1)) | ~ spl24_118), inference(forward_demodulation, [], [f393, f1181])).
fof(f7448, plain, (~ spl24_60 | ~ spl24_110), inference(avatar_split_clause, [], [f7447, f1145, f935])).
fof(f7447, plain, (~ (e4 = op(e2, e3)) | ~ spl24_110), inference(forward_demodulation, [], [f413, f1147])).
fof(f1147, plain, ((e4 = op(e0, e3)) | ~ spl24_110), inference(avatar_component_clause, [], [f1145])).
fof(f7443, plain, (~ spl24_35 | ~ spl24_45), inference(avatar_split_clause, [], [f7442, f872, f830])).
fof(f7442, plain, (~ (e4 = op(e3, e3)) | ~ spl24_45), inference(forward_demodulation, [], [f467, f874])).
fof(f467, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f7441, plain, (~ spl24_35 | ~ spl24_110), inference(avatar_split_clause, [], [f7440, f1145, f830])).
fof(f7440, plain, (~ (e4 = op(e3, e3)) | ~ spl24_110), inference(forward_demodulation, [], [f414, f1147])).
fof(f7439, plain, (~ spl24_30 | ~ spl24_45), inference(avatar_split_clause, [], [f7438, f872, f809])).
fof(f7438, plain, (~ (e4 = op(e3, e4)) | ~ spl24_45), inference(forward_demodulation, [], [f468, f874])).
fof(f7435, plain, (~ spl24_22 | ~ spl24_110 | ~ spl24_124 | spl24_204), inference(avatar_split_clause, [], [f7351, f1714, f1204, f1145, f776])).
fof(f7351, plain, (~ (e1 = op(e4, e0)) | (~ spl24_110 | ~ spl24_124 | spl24_204)), inference(backward_demodulation, [], [f7319, f1147])).
fof(f7319, plain, (~ (e1 = op(op(e0, e3), e0)) | (~ spl24_124 | spl24_204)), inference(forward_demodulation, [], [f1716, f1206])).
fof(f7406, plain, (~ spl24_43 | ~ spl24_48), inference(avatar_split_clause, [], [f7399, f885, f864])).
fof(f7399, plain, (~ (e2 = op(e3, e1)) | ~ spl24_48), inference(backward_demodulation, [], [f462, f887])).
fof(f462, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f4])).
fof(f7398, plain, (~ spl24_76 | ~ spl24_78), inference(avatar_contradiction_clause, [], [f7397])).
fof(f7397, plain, ($false | (~ spl24_76 | ~ spl24_78)), inference(subsumption_resolution, [], [f7396, f483])).
fof(f483, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f5])).
fof(f7396, plain, ((e0 = e2) | (~ spl24_76 | ~ spl24_78)), inference(backward_demodulation, [], [f1013, f1005])).
fof(f1005, plain, ((e0 = op(e1, e4)) | ~ spl24_76), inference(avatar_component_clause, [], [f1003])).
fof(f7384, plain, (spl24_45 | ~ spl24_94 | ~ spl24_213), inference(avatar_split_clause, [], [f7379, f1786, f1078, f872])).
fof(f7379, plain, ((e4 = op(e3, e1)) | (~ spl24_94 | ~ spl24_213)), inference(backward_demodulation, [], [f1787, f1080])).
fof(f1787, plain, ((e4 = op(op(e1, e1), e1)) | ~ spl24_213), inference(avatar_component_clause, [], [f1786])).
fof(f7383, plain, (~ spl24_43 | ~ spl24_94 | spl24_201), inference(avatar_split_clause, [], [f7377, f1690, f1078, f864])).
fof(f7377, plain, (~ (e2 = op(e3, e1)) | (~ spl24_94 | spl24_201)), inference(backward_demodulation, [], [f1692, f1080])).
fof(f1692, plain, (~ (e2 = op(op(e1, e1), e1)) | spl24_201), inference(avatar_component_clause, [], [f1690])).
fof(f7382, plain, (~ spl24_83 | ~ spl24_94 | spl24_182), inference(avatar_split_clause, [], [f7376, f1580, f1078, f1032])).
fof(f7376, plain, (~ (e2 = op(e1, e3)) | (~ spl24_94 | spl24_182)), inference(backward_demodulation, [], [f1582, f1080])).
fof(f7381, plain, (~ spl24_79 | ~ spl24_94), inference(avatar_split_clause, [], [f7370, f1078, f1015])).
fof(f7370, plain, (~ (e3 = op(e1, e4)) | ~ spl24_94), inference(backward_demodulation, [], [f448, f1080])).
fof(f7366, plain, (~ spl24_50 | ~ spl24_100), inference(avatar_split_clause, [], [f7362, f1103, f893])).
fof(f7362, plain, (~ (e4 = op(e3, e0)) | ~ spl24_100), inference(backward_demodulation, [], [f387, f1105])).
fof(f7353, plain, (~ spl24_10 | ~ spl24_110), inference(avatar_split_clause, [], [f7349, f1145, f725])).
fof(f7349, plain, (~ (e4 = op(e4, e3)) | ~ spl24_110), inference(backward_demodulation, [], [f415, f1147])).
fof(f415, plain, ~ (op(e0, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f7347, plain, (~ spl24_103 | ~ spl24_118), inference(avatar_split_clause, [], [f7343, f1179, f1116])).
fof(f7343, plain, (~ (e2 = op(e0, e4)) | ~ spl24_118), inference(backward_demodulation, [], [f438, f1181])).
fof(f7346, plain, (~ spl24_43 | ~ spl24_118), inference(avatar_split_clause, [], [f7340, f1179, f864])).
fof(f7340, plain, (~ (e2 = op(e3, e1)) | ~ spl24_118), inference(backward_demodulation, [], [f394, f1181])).
fof(f7337, plain, (spl24_48 | ~ spl24_31 | ~ spl24_167), inference(avatar_split_clause, [], [f7336, f1507, f814, f885])).
fof(f7336, plain, ((e2 = op(e3, e0)) | (~ spl24_31 | ~ spl24_167)), inference(forward_demodulation, [], [f1508, f816])).
fof(f816, plain, ((e0 = op(e3, e3)) | ~ spl24_31), inference(avatar_component_clause, [], [f814])).
fof(f1508, plain, ((e2 = op(e3, op(e3, e3))) | ~ spl24_167), inference(avatar_component_clause, [], [f1507])).
fof(f7333, plain, (~ spl24_107 | ~ spl24_31 | spl24_184), inference(avatar_split_clause, [], [f7332, f1596, f814, f1133])).
fof(f1596, plain, (spl24_184 <=> (e1 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_184])])).
fof(f7332, plain, (~ (e1 = op(e0, e3)) | (~ spl24_31 | spl24_184)), inference(forward_demodulation, [], [f1598, f816])).
fof(f1598, plain, (~ (e1 = op(op(e3, e3), e3)) | spl24_184), inference(avatar_component_clause, [], [f1596])).
fof(f7331, plain, (spl24_110 | ~ spl24_124 | ~ spl24_191), inference(avatar_split_clause, [], [f7330, f1632, f1204, f1145])).
fof(f7330, plain, ((e4 = op(e0, e3)) | (~ spl24_124 | ~ spl24_191)), inference(forward_demodulation, [], [f1633, f1206])).
fof(f7327, plain, (~ spl24_107 | ~ spl24_124 | spl24_206), inference(avatar_split_clause, [], [f7326, f1724, f1204, f1133])).
fof(f7326, plain, (~ (e1 = op(e0, e3)) | (~ spl24_124 | spl24_206)), inference(forward_demodulation, [], [f1726, f1206])).
fof(f7323, plain, (~ spl24_50 | ~ spl24_124 | spl24_214), inference(avatar_split_clause, [], [f7322, f1796, f1204, f893])).
fof(f7322, plain, (~ (e4 = op(e3, e0)) | (~ spl24_124 | spl24_214)), inference(forward_demodulation, [], [f1798, f1206])).
fof(f7311, plain, (~ spl24_109 | ~ spl24_124), inference(avatar_split_clause, [], [f7310, f1204, f1141])).
fof(f7310, plain, (~ (e3 = op(e0, e3)) | ~ spl24_124), inference(forward_demodulation, [], [f434, f1206])).
fof(f7283, plain, (~ spl24_20 | ~ spl24_15), inference(avatar_split_clause, [], [f7282, f746, f767])).
fof(f767, plain, (spl24_20 <=> (e4 = op(e4, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_20])])).
fof(f7282, plain, (~ (e4 = op(e4, e1)) | ~ spl24_15), inference(forward_demodulation, [], [f476, f748])).
fof(f7272, plain, (~ spl24_2 | ~ spl24_22), inference(avatar_split_clause, [], [f7271, f776, f692])).
fof(f7271, plain, (~ (e1 = op(e4, e4)) | ~ spl24_22), inference(forward_demodulation, [], [f475, f778])).
fof(f7270, plain, (~ spl24_2 | ~ spl24_27), inference(avatar_split_clause, [], [f7269, f797, f692])).
fof(f7269, plain, (~ (e1 = op(e4, e4)) | ~ spl24_27), inference(forward_demodulation, [], [f431, f799])).
fof(f799, plain, ((e1 = op(e3, e4)) | ~ spl24_27), inference(avatar_component_clause, [], [f797])).
fof(f7266, plain, (~ spl24_13 | ~ spl24_15), inference(avatar_contradiction_clause, [], [f7265])).
fof(f7265, plain, ($false | (~ spl24_13 | ~ spl24_15)), inference(subsumption_resolution, [], [f7264, f490])).
fof(f490, plain, ~ (e2 = e4), inference(cnf_transformation, [], [f5])).
fof(f7264, plain, ((e2 = e4) | (~ spl24_13 | ~ spl24_15)), inference(forward_demodulation, [], [f748, f740])).
fof(f740, plain, ((e2 = op(e4, e2)) | ~ spl24_13), inference(avatar_component_clause, [], [f738])).
fof(f7241, plain, (~ spl24_29 | ~ spl24_39), inference(avatar_split_clause, [], [f7239, f847, f805])).
fof(f7239, plain, (~ (e3 = op(e3, e4)) | ~ spl24_39), inference(backward_demodulation, [], [f470, f849])).
fof(f7231, plain, (~ spl24_35 | ~ spl24_50), inference(avatar_split_clause, [], [f7229, f893, f830])).
fof(f7229, plain, (~ (e4 = op(e3, e3)) | ~ spl24_50), inference(backward_demodulation, [], [f464, f895])).
fof(f7225, plain, (~ spl24_9 | ~ spl24_59), inference(avatar_split_clause, [], [f7222, f931, f721])).
fof(f7222, plain, (~ (e3 = op(e4, e3)) | ~ spl24_59), inference(backward_demodulation, [], [f420, f933])).
fof(f933, plain, ((e3 = op(e2, e3)) | ~ spl24_59), inference(avatar_component_clause, [], [f931])).
fof(f420, plain, ~ (op(e2, e3) = op(e4, e3)), inference(cnf_transformation, [], [f4])).
fof(f7214, plain, (~ spl24_53 | ~ spl24_63), inference(avatar_split_clause, [], [f7203, f948, f906])).
fof(f7203, plain, (~ (e2 = op(e2, e4)) | ~ spl24_63), inference(backward_demodulation, [], [f460, f950])).
fof(f7200, plain, (~ spl24_17 | ~ spl24_67), inference(avatar_split_clause, [], [f7193, f965, f755])).
fof(f7193, plain, (~ (e1 = op(e4, e1)) | ~ spl24_67), inference(backward_demodulation, [], [f400, f967])).
fof(f7172, plain, (~ spl24_65 | ~ spl24_90), inference(avatar_split_clause, [], [f7169, f1061, f956])).
fof(f7169, plain, (~ (e4 = op(e2, e2)) | ~ spl24_90), inference(backward_demodulation, [], [f406, f1063])).
fof(f7168, plain, (~ spl24_86 | ~ spl24_91), inference(avatar_split_clause, [], [f7157, f1066, f1045])).
fof(f7157, plain, (~ (e0 = op(e1, e2)) | ~ spl24_91), inference(backward_demodulation, [], [f446, f1068])).
fof(f7145, plain, (~ spl24_53 | ~ spl24_103), inference(avatar_split_clause, [], [f7141, f1116, f906])).
fof(f7141, plain, (~ (e2 = op(e2, e4)) | ~ spl24_103), inference(backward_demodulation, [], [f423, f1118])).
fof(f1118, plain, ((e2 = op(e0, e4)) | ~ spl24_103), inference(avatar_component_clause, [], [f1116])).
fof(f423, plain, ~ (op(e0, e4) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f7139, plain, (~ spl24_101 | ~ spl24_111), inference(avatar_split_clause, [], [f7133, f1150, f1108])).
fof(f7133, plain, (~ (e0 = op(e0, e4)) | ~ spl24_111), inference(backward_demodulation, [], [f440, f1152])).
fof(f7138, plain, (~ spl24_86 | ~ spl24_111), inference(avatar_split_clause, [], [f7130, f1150, f1045])).
fof(f7130, plain, (~ (e0 = op(e1, e2)) | ~ spl24_111), inference(backward_demodulation, [], [f402, f1152])).
fof(f7129, plain, (~ spl24_95 | ~ spl24_120), inference(avatar_split_clause, [], [f7123, f1187, f1082])).
fof(f7123, plain, (~ (e4 = op(e1, e1)) | ~ spl24_120), inference(backward_demodulation, [], [f392, f1189])).
fof(f7120, plain, (~ spl24_114 | ~ spl24_124), inference(avatar_split_clause, [], [f7111, f1204, f1162])).
fof(f7111, plain, (~ (e3 = op(e0, e2)) | ~ spl24_124), inference(backward_demodulation, [], [f433, f1206])).
fof(f7119, plain, (~ spl24_99 | ~ spl24_124), inference(avatar_split_clause, [], [f7107, f1204, f1099])).
fof(f7107, plain, (~ (e3 = op(e1, e0)) | ~ spl24_124), inference(backward_demodulation, [], [f382, f1206])).
fof(f7106, plain, (~ spl24_13 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f7105])).
fof(f7105, plain, ($false | (~ spl24_13 | ~ spl24_128)), inference(subsumption_resolution, [], [f7104, f490])).
fof(f7104, plain, ((e2 = e4) | (~ spl24_13 | ~ spl24_128)), inference(forward_demodulation, [], [f7092, f740])).
fof(f7092, plain, ((e4 = op(e4, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f330, f1223])).
fof(f7079, plain, (~ spl24_69 | spl24_174 | ~ spl24_182), inference(avatar_split_clause, [], [f7078, f1580, f1540, f973])).
fof(f7078, plain, (~ (e3 = op(e2, e1)) | (spl24_174 | ~ spl24_182)), inference(forward_demodulation, [], [f1542, f1581])).
fof(f1581, plain, ((e2 = op(e1, op(e1, e1))) | ~ spl24_182), inference(avatar_component_clause, [], [f1580])).
fof(f7070, plain, (spl24_103 | ~ spl24_1 | ~ spl24_199), inference(avatar_split_clause, [], [f7069, f1679, f688, f1116])).
fof(f1679, plain, (spl24_199 <=> (e2 = op(op(e4, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_199])])).
fof(f7069, plain, ((e2 = op(e0, e4)) | (~ spl24_1 | ~ spl24_199)), inference(forward_demodulation, [], [f1680, f690])).
fof(f1680, plain, ((e2 = op(op(e4, e4), e4)) | ~ spl24_199), inference(avatar_component_clause, [], [f1679])).
fof(f7068, plain, (~ spl24_98 | spl24_196 | ~ spl24_206), inference(avatar_split_clause, [], [f7065, f1724, f1658, f1095])).
fof(f1658, plain, (spl24_196 <=> (e2 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_196])])).
fof(f7065, plain, (~ (e2 = op(e1, e0)) | (spl24_196 | ~ spl24_206)), inference(backward_demodulation, [], [f1660, f1725])).
fof(f1725, plain, ((e1 = op(e0, op(e0, e0))) | ~ spl24_206), inference(avatar_component_clause, [], [f1724])).
fof(f1660, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | spl24_196), inference(avatar_component_clause, [], [f1658])).
fof(f7063, plain, (~ spl24_79 | ~ spl24_1 | ~ spl24_22 | spl24_161), inference(avatar_split_clause, [], [f7062, f1481, f776, f688, f1015])).
fof(f1481, plain, (spl24_161 <=> (e3 = op(op(e4, op(e4, e4)), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_161])])).
fof(f7062, plain, (~ (e3 = op(e1, e4)) | (~ spl24_1 | ~ spl24_22 | spl24_161)), inference(forward_demodulation, [], [f7061, f778])).
fof(f7061, plain, (~ (e3 = op(op(e4, e0), e4)) | (~ spl24_1 | spl24_161)), inference(forward_demodulation, [], [f1483, f690])).
fof(f1483, plain, (~ (e3 = op(op(e4, op(e4, e4)), e4)) | spl24_161), inference(avatar_component_clause, [], [f1481])).
fof(f7054, plain, (~ spl24_98 | ~ spl24_83), inference(avatar_split_clause, [], [f7053, f1032, f1095])).
fof(f7053, plain, (~ (e2 = op(e1, e0)) | ~ spl24_83), inference(forward_demodulation, [], [f444, f1034])).
fof(f7048, plain, (~ spl24_72 | ~ spl24_22), inference(avatar_split_clause, [], [f7047, f776, f986])).
fof(f7047, plain, (~ (e1 = op(e2, e0)) | ~ spl24_22), inference(forward_demodulation, [], [f390, f778])).
fof(f6987, plain, (~ spl24_28 | ~ spl24_53), inference(avatar_split_clause, [], [f6981, f906, f801])).
fof(f6981, plain, (~ (e2 = op(e3, e4)) | ~ spl24_53), inference(backward_demodulation, [], [f429, f908])).
fof(f908, plain, ((e2 = op(e2, e4)) | ~ spl24_53), inference(avatar_component_clause, [], [f906])).
fof(f6974, plain, (~ spl24_55 | ~ spl24_65), inference(avatar_split_clause, [], [f6963, f956, f914])).
fof(f6963, plain, (~ (e4 = op(e2, e4)) | ~ spl24_65), inference(backward_demodulation, [], [f460, f958])).
fof(f6973, plain, (~ spl24_15 | ~ spl24_65), inference(avatar_split_clause, [], [f6961, f956, f746])).
fof(f6961, plain, (~ (e4 = op(e4, e2)) | ~ spl24_65), inference(backward_demodulation, [], [f410, f958])).
fof(f6959, plain, (~ spl24_59 | ~ spl24_69), inference(avatar_split_clause, [], [f6955, f973, f931])).
fof(f6955, plain, (~ (e3 = op(e2, e3)) | ~ spl24_69), inference(backward_demodulation, [], [f457, f975])).
fof(f6958, plain, (~ spl24_19 | ~ spl24_69), inference(avatar_split_clause, [], [f6953, f973, f763])).
fof(f6953, plain, (~ (e3 = op(e4, e1)) | ~ spl24_69), inference(backward_demodulation, [], [f400, f975])).
fof(f6952, plain, (~ spl24_62 | ~ spl24_72), inference(avatar_split_clause, [], [f6943, f986, f944])).
fof(f6943, plain, (~ (e1 = op(e2, e2)) | ~ spl24_72), inference(backward_demodulation, [], [f453, f988])).
fof(f6951, plain, (~ spl24_67 | ~ spl24_72), inference(avatar_split_clause, [], [f6942, f986, f965])).
fof(f6942, plain, (~ (e1 = op(e2, e1)) | ~ spl24_72), inference(backward_demodulation, [], [f452, f988])).
fof(f452, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f4])).
fof(f6907, plain, (~ spl24_79 | ~ spl24_99), inference(avatar_split_clause, [], [f6905, f1099, f1015])).
fof(f6905, plain, (~ (e3 = op(e1, e4)) | ~ spl24_99), inference(backward_demodulation, [], [f445, f1101])).
fof(f6886, plain, (~ spl24_102 | ~ spl24_107), inference(avatar_split_clause, [], [f6881, f1133, f1112])).
fof(f6881, plain, (~ (e1 = op(e0, e4)) | ~ spl24_107), inference(backward_demodulation, [], [f441, f1135])).
fof(f6885, plain, (~ spl24_7 | ~ spl24_107), inference(avatar_split_clause, [], [f6880, f1133, f713])).
fof(f713, plain, (spl24_7 <=> (e1 = op(e4, e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_7])])).
fof(f6880, plain, (~ (e1 = op(e4, e3)) | ~ spl24_107), inference(backward_demodulation, [], [f415, f1135])).
fof(f6877, plain, (~ spl24_39 | ~ spl24_114), inference(avatar_split_clause, [], [f6872, f1162, f847])).
fof(f6872, plain, (~ (e3 = op(e3, e2)) | ~ spl24_114), inference(backward_demodulation, [], [f404, f1164])).
fof(f6868, plain, (~ spl24_93 | ~ spl24_118), inference(avatar_split_clause, [], [f6860, f1179, f1074])).
fof(f6860, plain, (~ (e2 = op(e1, e1)) | ~ spl24_118), inference(backward_demodulation, [], [f392, f1181])).
fof(f6856, plain, (~ spl24_100 | ~ spl24_125), inference(avatar_split_clause, [], [f6841, f1208, f1103])).
fof(f6841, plain, (~ (e4 = op(e1, e0)) | ~ spl24_125), inference(backward_demodulation, [], [f382, f1210])).
fof(f6840, plain, (~ spl24_1 | ~ spl24_130), inference(avatar_contradiction_clause, [], [f6839])).
fof(f6839, plain, ($false | (~ spl24_1 | ~ spl24_130)), inference(subsumption_resolution, [], [f6838, f485])).
fof(f485, plain, ~ (e0 = e4), inference(cnf_transformation, [], [f5])).
fof(f6838, plain, ((e0 = e4) | (~ spl24_1 | ~ spl24_130)), inference(forward_demodulation, [], [f6826, f690])).
fof(f6826, plain, ((e4 = op(e4, e4)) | ~ spl24_130), inference(backward_demodulation, [], [f330, f1231])).
fof(f6812, plain, (~ spl24_102 | ~ spl24_1 | spl24_183), inference(avatar_split_clause, [], [f6811, f1591, f688, f1112])).
fof(f1591, plain, (spl24_183 <=> (e1 = op(op(e4, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_183])])).
fof(f6811, plain, (~ (e1 = op(e0, e4)) | (~ spl24_1 | spl24_183)), inference(forward_demodulation, [], [f1593, f690])).
fof(f1593, plain, (~ (e1 = op(op(e4, e4), e4)) | spl24_183), inference(avatar_component_clause, [], [f1591])).
fof(f6810, plain, (~ spl24_7 | ~ spl24_35 | spl24_184), inference(avatar_split_clause, [], [f6809, f1596, f830, f713])).
fof(f6809, plain, (~ (e1 = op(e4, e3)) | (~ spl24_35 | spl24_184)), inference(forward_demodulation, [], [f1598, f832])).
fof(f6784, plain, (~ spl24_24 | ~ spl24_1 | spl24_157), inference(avatar_split_clause, [], [f6767, f1463, f688, f784])).
fof(f1463, plain, (spl24_157 <=> (e3 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_157])])).
fof(f6767, plain, (~ (e3 = op(e4, e0)) | (~ spl24_1 | spl24_157)), inference(backward_demodulation, [], [f1465, f690])).
fof(f1465, plain, (~ (e3 = op(e4, op(e4, e4))) | spl24_157), inference(avatar_component_clause, [], [f1463])).
fof(f6760, plain, (~ spl24_2 | ~ spl24_7), inference(avatar_split_clause, [], [f6756, f713, f692])).
fof(f6756, plain, (~ (e1 = op(e4, e4)) | ~ spl24_7), inference(backward_demodulation, [], [f481, f715])).
fof(f715, plain, ((e1 = op(e4, e3)) | ~ spl24_7), inference(avatar_component_clause, [], [f713])).
fof(f6752, plain, (~ spl24_26 | ~ spl24_35 | spl24_192), inference(avatar_split_clause, [], [f6748, f1637, f830, f793])).
fof(f6748, plain, (~ (e0 = op(e3, e4)) | (~ spl24_35 | spl24_192)), inference(backward_demodulation, [], [f1639, f832])).
fof(f6750, plain, (spl24_28 | ~ spl24_35 | ~ spl24_167), inference(avatar_split_clause, [], [f6745, f1507, f830, f801])).
fof(f6745, plain, ((e2 = op(e3, e4)) | (~ spl24_35 | ~ spl24_167)), inference(backward_demodulation, [], [f1508, f832])).
fof(f6741, plain, (~ spl24_26 | ~ spl24_41), inference(avatar_split_clause, [], [f6736, f856, f793])).
fof(f6736, plain, (~ (e0 = op(e3, e4)) | ~ spl24_41), inference(backward_demodulation, [], [f468, f858])).
fof(f6727, plain, (~ spl24_62 | ~ spl24_63), inference(avatar_contradiction_clause, [], [f6726])).
fof(f6726, plain, ($false | (~ spl24_62 | ~ spl24_63)), inference(subsumption_resolution, [], [f6725, f486])).
fof(f486, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f5])).
fof(f6725, plain, ((e1 = e2) | (~ spl24_62 | ~ spl24_63)), inference(backward_demodulation, [], [f950, f946])).
fof(f6688, plain, (~ spl24_2 | ~ spl24_102), inference(avatar_split_clause, [], [f6684, f1112, f692])).
fof(f6684, plain, (~ (e1 = op(e4, e4)) | ~ spl24_102), inference(backward_demodulation, [], [f425, f1114])).
fof(f6680, plain, (~ spl24_104 | ~ spl24_124), inference(avatar_split_clause, [], [f6674, f1204, f1120])).
fof(f6674, plain, (~ (e3 = op(e0, e4)) | ~ spl24_124), inference(backward_demodulation, [], [f435, f1206])).
fof(f6656, plain, (~ spl24_45 | spl24_177 | ~ spl24_178), inference(avatar_split_clause, [], [f6565, f1558, f1554, f872])).
fof(f1554, plain, (spl24_177 <=> (e4 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_177])])).
fof(f6565, plain, (~ (e4 = op(e3, e1)) | (spl24_177 | ~ spl24_178)), inference(backward_demodulation, [], [f1556, f1559])).
fof(f1559, plain, ((e3 = op(e1, op(e1, e1))) | ~ spl24_178), inference(avatar_component_clause, [], [f1558])).
fof(f1556, plain, (~ (e4 = op(op(e1, op(e1, e1)), e1)) | spl24_177), inference(avatar_component_clause, [], [f1554])).
fof(f6655, plain, (~ spl24_43 | ~ spl24_178 | spl24_180), inference(avatar_split_clause, [], [f6566, f1570, f1558, f864])).
fof(f6566, plain, (~ (e2 = op(e3, e1)) | (~ spl24_178 | spl24_180)), inference(backward_demodulation, [], [f1572, f1559])).
fof(f6654, plain, (~ spl24_2 | spl24_17 | ~ spl24_172), inference(avatar_contradiction_clause, [], [f6653])).
fof(f6653, plain, ($false | (~ spl24_2 | spl24_17 | ~ spl24_172)), inference(subsumption_resolution, [], [f6649, f756])).
fof(f756, plain, (~ (e1 = op(e4, e1)) | spl24_17), inference(avatar_component_clause, [], [f755])).
fof(f6649, plain, ((e1 = op(e4, e1)) | (~ spl24_2 | ~ spl24_172)), inference(backward_demodulation, [], [f1532, f694])).
fof(f694, plain, ((e1 = op(e4, e4)) | ~ spl24_2), inference(avatar_component_clause, [], [f692])).
fof(f1532, plain, ((e1 = op(e4, op(e4, e4))) | ~ spl24_172), inference(avatar_component_clause, [], [f1531])).
fof(f6639, plain, (~ spl24_4 | ~ spl24_19), inference(avatar_split_clause, [], [f6637, f763, f700])).
fof(f6637, plain, (~ (e3 = op(e4, e4)) | ~ spl24_19), inference(backward_demodulation, [], [f478, f765])).
fof(f6630, plain, (~ spl24_7 | ~ spl24_32), inference(avatar_split_clause, [], [f6617, f818, f713])).
fof(f6617, plain, (~ (e1 = op(e4, e3)) | ~ spl24_32), inference(backward_demodulation, [], [f421, f820])).
fof(f6614, plain, (~ spl24_35 | ~ spl24_85), inference(avatar_split_clause, [], [f6612, f1040, f830])).
fof(f6612, plain, (~ (e4 = op(e3, e3)) | ~ spl24_85), inference(backward_demodulation, [], [f417, f1042])).
fof(f6608, plain, (~ spl24_16 | ~ spl24_91), inference(avatar_split_clause, [], [f6597, f1066, f751])).
fof(f6597, plain, (~ (e0 = op(e4, e1)) | ~ spl24_91), inference(backward_demodulation, [], [f398, f1068])).
fof(f6584, plain, (~ spl24_102 | ~ spl24_122), inference(avatar_split_clause, [], [f6576, f1196, f1112])).
fof(f6576, plain, (~ (e1 = op(e0, e4)) | ~ spl24_122), inference(backward_demodulation, [], [f435, f1198])).
fof(f6583, plain, (~ spl24_47 | ~ spl24_122), inference(avatar_split_clause, [], [f6575, f1196, f881])).
fof(f6575, plain, (~ (e1 = op(e3, e0)) | ~ spl24_122), inference(backward_demodulation, [], [f384, f1198])).
fof(f6549, plain, (~ spl24_49 | ~ spl24_39), inference(avatar_split_clause, [], [f6548, f847, f889])).
fof(f6548, plain, (~ (e3 = op(e3, e0)) | ~ spl24_39), inference(forward_demodulation, [], [f463, f849])).
fof(f463, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f4])).
fof(f6535, plain, (~ spl24_6 | ~ spl24_16), inference(avatar_split_clause, [], [f6529, f751, f709])).
fof(f6529, plain, (~ (e0 = op(e4, e3)) | ~ spl24_16), inference(backward_demodulation, [], [f477, f753])).
fof(f6514, plain, (~ spl24_33 | ~ spl24_43), inference(avatar_split_clause, [], [f6510, f864, f822])).
fof(f6510, plain, (~ (e2 = op(e3, e3)) | ~ spl24_43), inference(backward_demodulation, [], [f467, f866])).
fof(f6503, plain, (~ spl24_81 | ~ spl24_94 | ~ spl24_120 | spl24_177), inference(avatar_contradiction_clause, [], [f6502])).
fof(f6502, plain, ($false | (~ spl24_81 | ~ spl24_94 | ~ spl24_120 | spl24_177)), inference(subsumption_resolution, [], [f6498, f1189])).
fof(f6498, plain, (~ (e4 = op(e0, e1)) | (~ spl24_81 | ~ spl24_94 | spl24_177)), inference(backward_demodulation, [], [f6481, f1026])).
fof(f6481, plain, (~ (e4 = op(op(e1, e3), e1)) | (~ spl24_94 | spl24_177)), inference(backward_demodulation, [], [f1556, f1080])).
fof(f6490, plain, (~ spl24_81 | ~ spl24_94 | spl24_205), inference(avatar_split_clause, [], [f6485, f1719, f1078, f1024])).
fof(f6485, plain, (~ (e0 = op(e1, e3)) | (~ spl24_94 | spl24_205)), inference(backward_demodulation, [], [f1721, f1080])).
fof(f6488, plain, (~ spl24_85 | ~ spl24_94 | spl24_175), inference(avatar_split_clause, [], [f6480, f1544, f1078, f1040])).
fof(f6480, plain, (~ (e4 = op(e1, e3)) | (~ spl24_94 | spl24_175)), inference(backward_demodulation, [], [f1546, f1080])).
fof(f6487, plain, (~ spl24_41 | ~ spl24_94 | spl24_173), inference(avatar_split_clause, [], [f6479, f1536, f1078, f856])).
fof(f6479, plain, (~ (e0 = op(e3, e1)) | (~ spl24_94 | spl24_173)), inference(backward_demodulation, [], [f1538, f1080])).
fof(f6434, plain, (~ spl24_78 | spl24_156 | ~ spl24_172), inference(avatar_split_clause, [], [f6429, f1531, f1459, f1011])).
fof(f1459, plain, (spl24_156 <=> (e2 = op(op(e4, op(e4, e4)), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_156])])).
fof(f6429, plain, (~ (e2 = op(e1, e4)) | (spl24_156 | ~ spl24_172)), inference(backward_demodulation, [], [f1461, f1532])).
fof(f1461, plain, (~ (e2 = op(op(e4, op(e4, e4)), e4)) | spl24_156), inference(avatar_component_clause, [], [f1459])).
fof(f6425, plain, (~ spl24_107 | spl24_171 | ~ spl24_192), inference(avatar_split_clause, [], [f6421, f1637, f1526, f1133])).
fof(f1526, plain, (spl24_171 <=> (e1 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_171])])).
fof(f6421, plain, (~ (e1 = op(e0, e3)) | (spl24_171 | ~ spl24_192)), inference(backward_demodulation, [], [f1528, f1638])).
fof(f1638, plain, ((e0 = op(e3, op(e3, e3))) | ~ spl24_192), inference(avatar_component_clause, [], [f1637])).
fof(f1528, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | spl24_171), inference(avatar_component_clause, [], [f1526])).
fof(f6424, plain, (~ spl24_108 | spl24_159 | ~ spl24_192), inference(avatar_split_clause, [], [f6419, f1637, f1472, f1137])).
fof(f1472, plain, (spl24_159 <=> (e2 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_159])])).
fof(f6419, plain, (~ (e2 = op(e0, e3)) | (spl24_159 | ~ spl24_192)), inference(backward_demodulation, [], [f1474, f1638])).
fof(f1474, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | spl24_159), inference(avatar_component_clause, [], [f1472])).
fof(f6411, plain, (~ spl24_80 | ~ spl24_55), inference(avatar_split_clause, [], [f6261, f914, f1019])).
fof(f1019, plain, (spl24_80 <=> (e4 = op(e1, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_80])])).
fof(f6261, plain, (~ (e4 = op(e1, e4)) | ~ spl24_55), inference(forward_demodulation, [], [f426, f916])).
fof(f426, plain, ~ (op(e1, e4) = op(e2, e4)), inference(cnf_transformation, [], [f4])).
fof(f6388, plain, (~ spl24_29 | ~ spl24_30), inference(avatar_contradiction_clause, [], [f6387])).
fof(f6387, plain, ($false | (~ spl24_29 | ~ spl24_30)), inference(subsumption_resolution, [], [f6386, f491])).
fof(f6386, plain, ((e3 = e4) | (~ spl24_29 | ~ spl24_30)), inference(backward_demodulation, [], [f811, f807])).
fof(f811, plain, ((e4 = op(e3, e4)) | ~ spl24_30), inference(avatar_component_clause, [], [f809])).
fof(f6356, plain, (~ spl24_63 | ~ spl24_64), inference(avatar_contradiction_clause, [], [f6355])).
fof(f6355, plain, ($false | (~ spl24_63 | ~ spl24_64)), inference(subsumption_resolution, [], [f6354, f489])).
fof(f6354, plain, ((e2 = e3) | (~ spl24_63 | ~ spl24_64)), inference(forward_demodulation, [], [f954, f950])).
fof(f6350, plain, (~ spl24_93 | ~ spl24_95), inference(avatar_contradiction_clause, [], [f6349])).
fof(f6349, plain, ($false | (~ spl24_93 | ~ spl24_95)), inference(subsumption_resolution, [], [f6348, f490])).
fof(f6348, plain, ((e2 = e4) | (~ spl24_93 | ~ spl24_95)), inference(backward_demodulation, [], [f1084, f1076])).
fof(f6346, plain, (~ spl24_19 | ~ spl24_95 | spl24_209), inference(avatar_split_clause, [], [f6342, f1746, f1082, f763])).
fof(f6342, plain, (~ (e3 = op(e4, e1)) | (~ spl24_95 | spl24_209)), inference(backward_demodulation, [], [f1748, f1084])).
fof(f6345, plain, (~ spl24_16 | ~ spl24_95 | spl24_173), inference(avatar_split_clause, [], [f6339, f1536, f1082, f751])).
fof(f6339, plain, (~ (e0 = op(e4, e1)) | (~ spl24_95 | spl24_173)), inference(backward_demodulation, [], [f1538, f1084])).
fof(f6331, plain, (~ spl24_19 | ~ spl24_119), inference(avatar_split_clause, [], [f6328, f1183, f763])).
fof(f6328, plain, (~ (e3 = op(e4, e1)) | ~ spl24_119), inference(backward_demodulation, [], [f395, f1185])).
fof(f6298, plain, (~ spl24_120 | spl24_177 | ~ spl24_205), inference(avatar_split_clause, [], [f6296, f1719, f1554, f1187])).
fof(f6296, plain, (~ (e4 = op(e0, e1)) | (spl24_177 | ~ spl24_205)), inference(backward_demodulation, [], [f1556, f1720])).
fof(f6250, plain, (~ spl24_26 | ~ spl24_76), inference(avatar_split_clause, [], [f6249, f1003, f793])).
fof(f6249, plain, (~ (e0 = op(e3, e4)) | ~ spl24_76), inference(forward_demodulation, [], [f427, f1005])).
fof(f6248, plain, (~ spl24_28 | ~ spl24_103), inference(avatar_split_clause, [], [f6247, f1116, f801])).
fof(f6247, plain, (~ (e2 = op(e3, e4)) | ~ spl24_103), inference(forward_demodulation, [], [f424, f1118])).
fof(f6233, plain, (~ spl24_16 | ~ spl24_17), inference(avatar_contradiction_clause, [], [f6232])).
fof(f6232, plain, ($false | (~ spl24_16 | ~ spl24_17)), inference(subsumption_resolution, [], [f6231, f482])).
fof(f482, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f5])).
fof(f6231, plain, ((e0 = e1) | (~ spl24_16 | ~ spl24_17)), inference(backward_demodulation, [], [f757, f753])).
fof(f6223, plain, (~ spl24_3 | ~ spl24_23), inference(avatar_split_clause, [], [f6220, f780, f696])).
fof(f6220, plain, (~ (e2 = op(e4, e4)) | ~ spl24_23), inference(backward_demodulation, [], [f475, f782])).
fof(f6168, plain, (~ spl24_76 | ~ spl24_77), inference(avatar_contradiction_clause, [], [f6167])).
fof(f6167, plain, ($false | (~ spl24_76 | ~ spl24_77)), inference(subsumption_resolution, [], [f6166, f482])).
fof(f6166, plain, ((e0 = e1) | (~ spl24_76 | ~ spl24_77)), inference(backward_demodulation, [], [f1009, f1005])).
fof(f6159, plain, (~ spl24_33 | ~ spl24_83), inference(avatar_split_clause, [], [f6154, f1032, f822])).
fof(f6154, plain, (~ (e2 = op(e3, e3)) | ~ spl24_83), inference(backward_demodulation, [], [f417, f1034])).
fof(f6141, plain, (spl24_83 | ~ spl24_94 | ~ spl24_182), inference(avatar_split_clause, [], [f6134, f1580, f1078, f1032])).
fof(f6134, plain, ((e2 = op(e1, e3)) | (~ spl24_94 | ~ spl24_182)), inference(backward_demodulation, [], [f1581, f1080])).
fof(f6140, plain, (~ spl24_84 | ~ spl24_94), inference(avatar_split_clause, [], [f6130, f1078, f1036])).
fof(f6130, plain, (~ (e3 = op(e1, e3)) | ~ spl24_94), inference(backward_demodulation, [], [f447, f1080])).
fof(f6139, plain, (~ spl24_44 | ~ spl24_94), inference(avatar_split_clause, [], [f6127, f1078, f868])).
fof(f868, plain, (spl24_44 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_44])])).
fof(f6127, plain, (~ (e3 = op(e3, e1)) | ~ spl24_94), inference(backward_demodulation, [], [f397, f1080])).
fof(f6117, plain, (~ spl24_3 | ~ spl24_103), inference(avatar_split_clause, [], [f6114, f1116, f696])).
fof(f6114, plain, (~ (e2 = op(e4, e4)) | ~ spl24_103), inference(backward_demodulation, [], [f425, f1118])).
fof(f6091, plain, (~ spl24_110 | ~ spl24_120), inference(avatar_split_clause, [], [f6088, f1187, f1145])).
fof(f6088, plain, (~ (e4 = op(e0, e3)) | ~ spl24_120), inference(backward_demodulation, [], [f437, f1189])).
fof(f6090, plain, (~ spl24_20 | ~ spl24_120), inference(avatar_split_clause, [], [f6086, f1187, f767])).
fof(f6086, plain, (~ (e4 = op(e4, e1)) | ~ spl24_120), inference(backward_demodulation, [], [f395, f1189])).
fof(f6047, plain, (~ spl24_155 | ~ spl24_199), inference(avatar_contradiction_clause, [], [f6046])).
fof(f6046, plain, ($false | (~ spl24_155 | ~ spl24_199)), inference(subsumption_resolution, [], [f6045, f483])).
fof(f6045, plain, ((e0 = e2) | (~ spl24_155 | ~ spl24_199)), inference(backward_demodulation, [], [f1680, f1456])).
fof(f1456, plain, ((e0 = op(op(e4, e4), e4)) | ~ spl24_155), inference(avatar_component_clause, [], [f1455])).
fof(f1455, plain, (spl24_155 <=> (e0 = op(op(e4, e4), e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_155])])).
fof(f6032, plain, (~ spl24_110 | spl24_166 | ~ spl24_192), inference(avatar_split_clause, [], [f6030, f1637, f1503, f1145])).
fof(f6030, plain, (~ (e4 = op(e0, e3)) | (spl24_166 | ~ spl24_192)), inference(backward_demodulation, [], [f1505, f1638])).
fof(f1505, plain, (~ (e4 = op(op(e3, op(e3, e3)), e3)) | spl24_166), inference(avatar_component_clause, [], [f1503])).
fof(f6023, plain, (~ spl24_104 | spl24_161 | ~ spl24_188), inference(avatar_split_clause, [], [f6022, f1619, f1481, f1120])).
fof(f1619, plain, (spl24_188 <=> (e0 = op(e4, op(e4, e4)))), introduced(avatar_definition, [new_symbols(naming, [spl24_188])])).
fof(f6022, plain, (~ (e3 = op(e0, e4)) | (spl24_161 | ~ spl24_188)), inference(forward_demodulation, [], [f1483, f1620])).
fof(f1620, plain, ((e0 = op(e4, op(e4, e4))) | ~ spl24_188), inference(avatar_component_clause, [], [f1619])).
fof(f6021, plain, (~ spl24_103 | spl24_156 | ~ spl24_188), inference(avatar_split_clause, [], [f6020, f1619, f1459, f1116])).
fof(f6020, plain, (~ (e2 = op(e0, e4)) | (spl24_156 | ~ spl24_188)), inference(forward_demodulation, [], [f1461, f1620])).
fof(f6019, plain, (~ spl24_115 | spl24_168 | ~ spl24_197), inference(avatar_split_clause, [], [f5838, f1663, f1512, f1166])).
fof(f5838, plain, (~ (e4 = op(e0, e2)) | (spl24_168 | ~ spl24_197)), inference(backward_demodulation, [], [f1514, f1664])).
fof(f1664, plain, ((e0 = op(e2, op(e2, e2))) | ~ spl24_197), inference(avatar_component_clause, [], [f1663])).
fof(f6000, plain, (~ spl24_1 | ~ spl24_11), inference(avatar_split_clause, [], [f5995, f730, f688])).
fof(f5995, plain, (~ (e0 = op(e4, e4)) | ~ spl24_11), inference(backward_demodulation, [], [f480, f732])).
fof(f732, plain, ((e0 = op(e4, e2)) | ~ spl24_11), inference(avatar_component_clause, [], [f730])).
fof(f480, plain, ~ (op(e4, e2) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f5994, plain, (~ spl24_14 | ~ spl24_24), inference(avatar_split_clause, [], [f5991, f784, f742])).
fof(f5991, plain, (~ (e3 = op(e4, e2)) | ~ spl24_24), inference(backward_demodulation, [], [f473, f786])).
fof(f5984, plain, (~ spl24_29 | ~ spl24_104), inference(avatar_split_clause, [], [f5981, f1120, f805])).
fof(f5981, plain, (~ (e3 = op(e3, e4)) | ~ spl24_104), inference(backward_demodulation, [], [f424, f1122])).
fof(f5977, plain, (~ spl24_23 | ~ spl24_123), inference(avatar_split_clause, [], [f5971, f1200, f780])).
fof(f5971, plain, (~ (e2 = op(e4, e0)) | ~ spl24_123), inference(backward_demodulation, [], [f385, f1202])).
fof(f5966, plain, (~ spl24_183 | ~ spl24_199), inference(avatar_contradiction_clause, [], [f5965])).
fof(f5965, plain, ($false | (~ spl24_183 | ~ spl24_199)), inference(subsumption_resolution, [], [f5964, f486])).
fof(f5964, plain, ((e1 = e2) | (~ spl24_183 | ~ spl24_199)), inference(backward_demodulation, [], [f1680, f1592])).
fof(f1592, plain, ((e1 = op(op(e4, e4), e4)) | ~ spl24_183), inference(avatar_component_clause, [], [f1591])).
fof(f5950, plain, (~ spl24_40 | ~ spl24_64 | spl24_212), inference(avatar_contradiction_clause, [], [f5949])).
fof(f5949, plain, ($false | (~ spl24_40 | ~ spl24_64 | spl24_212)), inference(subsumption_resolution, [], [f5948, f853])).
fof(f5948, plain, (~ (e4 = op(e3, e2)) | (~ spl24_64 | spl24_212)), inference(forward_demodulation, [], [f1782, f954])).
fof(f5936, plain, (~ spl24_105 | ~ spl24_110), inference(avatar_split_clause, [], [f5935, f1145, f1124])).
fof(f5935, plain, (~ (e4 = op(e0, e4)) | ~ spl24_110), inference(forward_demodulation, [], [f441, f1147])).
fof(f5925, plain, (~ spl24_14 | ~ spl24_64), inference(avatar_split_clause, [], [f5924, f952, f742])).
fof(f5924, plain, (~ (e3 = op(e4, e2)) | ~ spl24_64), inference(forward_demodulation, [], [f410, f954])).
fof(f5919, plain, (~ spl24_5 | ~ spl24_20), inference(avatar_split_clause, [], [f5918, f767, f704])).
fof(f5918, plain, (~ (e4 = op(e4, e4)) | ~ spl24_20), inference(forward_demodulation, [], [f478, f769])).
fof(f769, plain, ((e4 = op(e4, e1)) | ~ spl24_20), inference(avatar_component_clause, [], [f767])).
fof(f5914, plain, (~ spl24_2 | ~ spl24_52), inference(avatar_split_clause, [], [f5913, f902, f692])).
fof(f5913, plain, (~ (e1 = op(e4, e4)) | ~ spl24_52), inference(forward_demodulation, [], [f430, f904])).
fof(f430, plain, ~ (op(e2, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f5909, plain, (~ spl24_8 | ~ spl24_23), inference(avatar_split_clause, [], [f5904, f780, f717])).
fof(f5904, plain, (~ (e2 = op(e4, e3)) | ~ spl24_23), inference(backward_demodulation, [], [f474, f782])).
fof(f5886, plain, (~ spl24_36 | ~ spl24_40), inference(avatar_contradiction_clause, [], [f5885])).
fof(f5885, plain, ($false | (~ spl24_36 | ~ spl24_40)), inference(subsumption_resolution, [], [f5882, f485])).
fof(f5882, plain, ((e0 = e4) | (~ spl24_36 | ~ spl24_40)), inference(backward_demodulation, [], [f853, f837])).
fof(f5867, plain, (~ spl24_31 | ~ spl24_56), inference(avatar_split_clause, [], [f5860, f919, f814])).
fof(f5860, plain, (~ (e0 = op(e3, e3)) | ~ spl24_56), inference(backward_demodulation, [], [f419, f921])).
fof(f419, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f4])).
fof(f5859, plain, (spl24_56 | ~ spl24_64 | ~ spl24_197), inference(avatar_split_clause, [], [f5855, f1663, f952, f919])).
fof(f5855, plain, ((e0 = op(e2, e3)) | (~ spl24_64 | ~ spl24_197)), inference(backward_demodulation, [], [f1664, f954])).
fof(f5857, plain, (~ spl24_54 | ~ spl24_64), inference(avatar_split_clause, [], [f5852, f952, f910])).
fof(f5852, plain, (~ (e3 = op(e2, e4)) | ~ spl24_64), inference(backward_demodulation, [], [f460, f954])).
fof(f5850, plain, (~ spl24_22 | ~ spl24_1 | spl24_172), inference(avatar_split_clause, [], [f5849, f1531, f688, f776])).
fof(f5849, plain, (~ (e1 = op(e4, e0)) | (~ spl24_1 | spl24_172)), inference(forward_demodulation, [], [f1533, f690])).
fof(f5846, plain, (~ spl24_110 | ~ spl24_124 | spl24_191), inference(avatar_contradiction_clause, [], [f5845])).
fof(f5845, plain, ($false | (~ spl24_110 | ~ spl24_124 | spl24_191)), inference(subsumption_resolution, [], [f5844, f1147])).
fof(f5844, plain, (~ (e4 = op(e0, e3)) | (~ spl24_124 | spl24_191)), inference(forward_demodulation, [], [f1634, f1206])).
fof(f5842, plain, (~ spl24_112 | spl24_179 | ~ spl24_197), inference(avatar_contradiction_clause, [], [f5841])).
fof(f5841, plain, ($false | (~ spl24_112 | spl24_179 | ~ spl24_197)), inference(subsumption_resolution, [], [f5835, f1156])).
fof(f5835, plain, (~ (e1 = op(e0, e2)) | (spl24_179 | ~ spl24_197)), inference(backward_demodulation, [], [f1566, f1664])).
fof(f1566, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | spl24_179), inference(avatar_component_clause, [], [f1564])).
fof(f1564, plain, (spl24_179 <=> (e1 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_179])])).
fof(f5829, plain, (~ spl24_23 | ~ spl24_110 | ~ spl24_124 | spl24_196), inference(avatar_split_clause, [], [f5828, f1658, f1204, f1145, f780])).
fof(f5828, plain, (~ (e2 = op(e4, e0)) | (~ spl24_110 | ~ spl24_124 | spl24_196)), inference(forward_demodulation, [], [f5827, f1147])).
fof(f5827, plain, (~ (e2 = op(op(e0, e3), e0)) | (~ spl24_124 | spl24_196)), inference(forward_demodulation, [], [f1660, f1206])).
fof(f5818, plain, (~ spl24_65 | ~ spl24_40), inference(avatar_split_clause, [], [f5118, f851, f956])).
fof(f5118, plain, (~ (e4 = op(e2, e2)) | ~ spl24_40), inference(forward_demodulation, [], [f409, f853])).
fof(f5806, plain, (~ spl24_50 | ~ spl24_75), inference(avatar_split_clause, [], [f5805, f998, f893])).
fof(f5805, plain, (~ (e4 = op(e3, e0)) | ~ spl24_75), inference(forward_demodulation, [], [f389, f1000])).
fof(f5804, plain, (~ spl24_49 | ~ spl24_124), inference(avatar_split_clause, [], [f5803, f1204, f889])).
fof(f5803, plain, (~ (e3 = op(e3, e0)) | ~ spl24_124), inference(forward_demodulation, [], [f384, f1206])).
fof(f5798, plain, (~ spl24_29 | ~ spl24_44), inference(avatar_split_clause, [], [f5612, f868, f805])).
fof(f5612, plain, (~ (e3 = op(e3, e4)) | ~ spl24_44), inference(backward_demodulation, [], [f468, f870])).
fof(f870, plain, ((e3 = op(e3, e1)) | ~ spl24_44), inference(avatar_component_clause, [], [f868])).
fof(f5796, plain, (~ spl24_9 | ~ spl24_14), inference(avatar_split_clause, [], [f5795, f742, f721])).
fof(f5795, plain, (~ (e3 = op(e4, e3)) | ~ spl24_14), inference(forward_demodulation, [], [f479, f744])).
fof(f744, plain, ((e3 = op(e4, e2)) | ~ spl24_14), inference(avatar_component_clause, [], [f742])).
fof(f5779, plain, (~ spl24_12 | ~ spl24_61 | ~ spl24_75 | spl24_179), inference(avatar_split_clause, [], [f5774, f1564, f998, f940, f734])).
fof(f5774, plain, (~ (e1 = op(e4, e2)) | (~ spl24_61 | ~ spl24_75 | spl24_179)), inference(backward_demodulation, [], [f5103, f1000])).
fof(f5103, plain, (~ (e1 = op(op(e2, e0), e2)) | (~ spl24_61 | spl24_179)), inference(forward_demodulation, [], [f1566, f942])).
fof(f5778, plain, (~ spl24_14 | ~ spl24_61 | ~ spl24_75 | spl24_164), inference(avatar_split_clause, [], [f5773, f1494, f998, f940, f742])).
fof(f5773, plain, (~ (e3 = op(e4, e2)) | (~ spl24_61 | ~ spl24_75 | spl24_164)), inference(backward_demodulation, [], [f5100, f1000])).
fof(f5100, plain, (~ (e3 = op(op(e2, e0), e2)) | (~ spl24_61 | spl24_164)), inference(forward_demodulation, [], [f1496, f942])).
fof(f5767, plain, (~ spl24_12 | ~ spl24_112), inference(avatar_split_clause, [], [f5762, f1154, f734])).
fof(f5762, plain, (~ (e1 = op(e4, e2)) | ~ spl24_112), inference(backward_demodulation, [], [f405, f1156])).
fof(f5730, plain, (~ spl24_72 | ~ spl24_61 | spl24_181), inference(avatar_split_clause, [], [f5729, f1575, f940, f986])).
fof(f5729, plain, (~ (e1 = op(e2, e0)) | (~ spl24_61 | spl24_181)), inference(forward_demodulation, [], [f1577, f942])).
fof(f1577, plain, (~ (e1 = op(e2, op(e2, e2))) | spl24_181), inference(avatar_component_clause, [], [f1575])).
fof(f5725, plain, (~ spl24_114 | ~ spl24_61 | spl24_208), inference(avatar_split_clause, [], [f5724, f1740, f940, f1162])).
fof(f5724, plain, (~ (e3 = op(e0, e2)) | (~ spl24_61 | spl24_208)), inference(forward_demodulation, [], [f1742, f942])).
fof(f5717, plain, (~ spl24_75 | spl24_193 | ~ spl24_198), inference(avatar_split_clause, [], [f5716, f1668, f1642, f998])).
fof(f5716, plain, (~ (e4 = op(e2, e0)) | (spl24_193 | ~ spl24_198)), inference(forward_demodulation, [], [f1644, f1669])).
fof(f5700, plain, (~ spl24_73 | ~ spl24_68), inference(avatar_split_clause, [], [f5699, f969, f990])).
fof(f5699, plain, (~ (e2 = op(e2, e0)) | ~ spl24_68), inference(forward_demodulation, [], [f452, f971])).
fof(f971, plain, ((e2 = op(e2, e1)) | ~ spl24_68), inference(avatar_component_clause, [], [f969])).
fof(f5697, plain, (~ spl24_74 | ~ spl24_61 | spl24_169), inference(avatar_split_clause, [], [f5476, f1516, f940, f994])).
fof(f5476, plain, (~ (e3 = op(e2, e0)) | (~ spl24_61 | spl24_169)), inference(forward_demodulation, [], [f1518, f942])).
fof(f1518, plain, (~ (e3 = op(e2, op(e2, e2))) | spl24_169), inference(avatar_component_clause, [], [f1516])).
fof(f5694, plain, (~ spl24_60 | ~ spl24_31 | ~ spl24_48 | spl24_166), inference(avatar_split_clause, [], [f5652, f1503, f885, f814, f935])).
fof(f5652, plain, (~ (e4 = op(e2, e3)) | (~ spl24_31 | ~ spl24_48 | spl24_166)), inference(forward_demodulation, [], [f5651, f887])).
fof(f5651, plain, (~ (e4 = op(op(e3, e0), e3)) | (~ spl24_31 | spl24_166)), inference(forward_demodulation, [], [f1505, f816])).
fof(f5685, plain, (~ spl24_13 | ~ spl24_88), inference(avatar_split_clause, [], [f5684, f1053, f738])).
fof(f5684, plain, (~ (e2 = op(e4, e2)) | ~ spl24_88), inference(forward_demodulation, [], [f408, f1055])).
fof(f1055, plain, ((e2 = op(e1, e2)) | ~ spl24_88), inference(avatar_component_clause, [], [f1053])).
fof(f5672, plain, (~ spl24_24 | ~ spl24_125 | spl24_210), inference(avatar_contradiction_clause, [], [f5671])).
fof(f5671, plain, ($false | (~ spl24_24 | ~ spl24_125 | spl24_210)), inference(subsumption_resolution, [], [f5670, f786])).
fof(f5670, plain, (~ (e3 = op(e4, e0)) | (~ spl24_125 | spl24_210)), inference(forward_demodulation, [], [f1758, f1210])).
fof(f5669, plain, (~ spl24_1 | ~ spl24_103 | spl24_199), inference(avatar_contradiction_clause, [], [f5668])).
fof(f5668, plain, ($false | (~ spl24_1 | ~ spl24_103 | spl24_199)), inference(subsumption_resolution, [], [f5667, f1118])).
fof(f5667, plain, (~ (e2 = op(e0, e4)) | (~ spl24_1 | spl24_199)), inference(forward_demodulation, [], [f1681, f690])).
fof(f1681, plain, (~ (e2 = op(op(e4, e4), e4)) | spl24_199), inference(avatar_component_clause, [], [f1679])).
fof(f5650, plain, (~ spl24_31 | ~ spl24_48 | spl24_167), inference(avatar_contradiction_clause, [], [f5649])).
fof(f5649, plain, ($false | (~ spl24_31 | ~ spl24_48 | spl24_167)), inference(subsumption_resolution, [], [f5648, f887])).
fof(f5648, plain, (~ (e2 = op(e3, e0)) | (~ spl24_31 | spl24_167)), inference(forward_demodulation, [], [f1509, f816])).
fof(f5643, plain, (~ spl24_10 | ~ spl24_20), inference(avatar_split_clause, [], [f5642, f767, f725])).
fof(f5642, plain, (~ (e4 = op(e4, e3)) | ~ spl24_20), inference(backward_demodulation, [], [f477, f769])).
fof(f5641, plain, (~ spl24_1 | ~ spl24_24 | ~ spl24_27 | spl24_170), inference(avatar_contradiction_clause, [], [f5640])).
fof(f5640, plain, ($false | (~ spl24_1 | ~ spl24_24 | ~ spl24_27 | spl24_170)), inference(subsumption_resolution, [], [f5635, f799])).
fof(f5635, plain, (~ (e1 = op(e3, e4)) | (~ spl24_1 | ~ spl24_24 | spl24_170)), inference(backward_demodulation, [], [f5425, f786])).
fof(f5425, plain, (~ (e1 = op(op(e4, e0), e4)) | (~ spl24_1 | spl24_170)), inference(backward_demodulation, [], [f1523, f690])).
fof(f5637, plain, (~ spl24_19 | ~ spl24_24), inference(avatar_split_clause, [], [f5630, f784, f763])).
fof(f5630, plain, (~ (e3 = op(e4, e1)) | ~ spl24_24), inference(backward_demodulation, [], [f472, f786])).
fof(f5615, plain, (~ spl24_34 | ~ spl24_44), inference(avatar_split_clause, [], [f5611, f868, f826])).
fof(f5611, plain, (~ (e3 = op(e3, e3)) | ~ spl24_44), inference(backward_demodulation, [], [f467, f870])).
fof(f5614, plain, (~ spl24_19 | ~ spl24_44), inference(avatar_split_clause, [], [f5610, f868, f763])).
fof(f5610, plain, (~ (e3 = op(e4, e1)) | ~ spl24_44), inference(backward_demodulation, [], [f401, f870])).
fof(f5599, plain, (~ spl24_58 | ~ spl24_68), inference(avatar_split_clause, [], [f5596, f969, f927])).
fof(f5596, plain, (~ (e2 = op(e2, e3)) | ~ spl24_68), inference(backward_demodulation, [], [f457, f971])).
fof(f5593, plain, (~ spl24_30 | ~ spl24_80), inference(avatar_split_clause, [], [f5592, f1019, f809])).
fof(f5592, plain, (~ (e4 = op(e3, e4)) | ~ spl24_80), inference(backward_demodulation, [], [f427, f1021])).
fof(f1021, plain, ((e4 = op(e1, e4)) | ~ spl24_80), inference(avatar_component_clause, [], [f1019])).
fof(f5569, plain, (~ spl24_82 | ~ spl24_92), inference(avatar_split_clause, [], [f5555, f1070, f1028])).
fof(f5555, plain, (~ (e1 = op(e1, e3)) | ~ spl24_92), inference(backward_demodulation, [], [f447, f1072])).
fof(f1072, plain, ((e1 = op(e1, e1)) | ~ spl24_92), inference(avatar_component_clause, [], [f1070])).
fof(f5550, plain, (~ spl24_91 | ~ spl24_96), inference(avatar_split_clause, [], [f5541, f1087, f1066])).
fof(f1087, plain, (spl24_96 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_96])])).
fof(f5541, plain, (~ (e0 = op(e1, e1)) | ~ spl24_96), inference(backward_demodulation, [], [f442, f1089])).
fof(f1089, plain, ((e0 = op(e1, e0)) | ~ spl24_96), inference(avatar_component_clause, [], [f1087])).
fof(f5549, plain, (~ spl24_46 | ~ spl24_96), inference(avatar_split_clause, [], [f5539, f1087, f877])).
fof(f5539, plain, (~ (e0 = op(e3, e0)) | ~ spl24_96), inference(backward_demodulation, [], [f387, f1089])).
fof(f5533, plain, (~ spl24_78 | ~ spl24_103), inference(avatar_split_clause, [], [f5525, f1116, f1011])).
fof(f5525, plain, (~ (e2 = op(e1, e4)) | ~ spl24_103), inference(backward_demodulation, [], [f422, f1118])).
fof(f5523, plain, (~ spl24_82 | ~ spl24_107), inference(avatar_split_clause, [], [f5515, f1133, f1028])).
fof(f5515, plain, (~ (e1 = op(e1, e3)) | ~ spl24_107), inference(backward_demodulation, [], [f412, f1135])).
fof(f5514, plain, (~ spl24_106 | ~ spl24_116), inference(avatar_split_clause, [], [f5507, f1171, f1129])).
fof(f1171, plain, (spl24_116 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl24_116])])).
fof(f5507, plain, (~ (e0 = op(e0, e3)) | ~ spl24_116), inference(backward_demodulation, [], [f437, f1173])).
fof(f1173, plain, ((e0 = op(e0, e1)) | ~ spl24_116), inference(avatar_component_clause, [], [f1171])).
fof(f5513, plain, (~ spl24_91 | ~ spl24_116), inference(avatar_split_clause, [], [f5503, f1171, f1066])).
fof(f5503, plain, (~ (e0 = op(e1, e1)) | ~ spl24_116), inference(backward_demodulation, [], [f392, f1173])).
fof(f5502, plain, (spl24_20 | ~ spl24_127), inference(avatar_split_clause, [], [f5489, f1217, f767])).
fof(f1217, plain, (spl24_127 <=> (e1 = unit)), introduced(avatar_definition, [new_symbols(naming, [spl24_127])])).
fof(f5489, plain, ((e4 = op(e4, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f330, f1219])).
fof(f1219, plain, ((e1 = unit) | ~ spl24_127), inference(avatar_component_clause, [], [f1217])).
fof(f5501, plain, (spl24_80 | ~ spl24_127), inference(avatar_split_clause, [], [f5488, f1217, f1019])).
fof(f5488, plain, ((e4 = op(e1, e4)) | ~ spl24_127), inference(backward_demodulation, [], [f329, f1219])).
fof(f5500, plain, (spl24_44 | ~ spl24_127), inference(avatar_split_clause, [], [f5487, f1217, f868])).
fof(f5487, plain, ((e3 = op(e3, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f328, f1219])).
fof(f5499, plain, (spl24_84 | ~ spl24_127), inference(avatar_split_clause, [], [f5486, f1217, f1036])).
fof(f5486, plain, ((e3 = op(e1, e3)) | ~ spl24_127), inference(backward_demodulation, [], [f327, f1219])).
fof(f5498, plain, (spl24_68 | ~ spl24_127), inference(avatar_split_clause, [], [f5485, f1217, f969])).
fof(f5485, plain, ((e2 = op(e2, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f326, f1219])).
fof(f5497, plain, (spl24_88 | ~ spl24_127), inference(avatar_split_clause, [], [f5484, f1217, f1053])).
fof(f5484, plain, ((e2 = op(e1, e2)) | ~ spl24_127), inference(backward_demodulation, [], [f325, f1219])).
fof(f5496, plain, (spl24_92 | ~ spl24_127), inference(avatar_split_clause, [], [f5483, f1217, f1070])).
fof(f5483, plain, ((e1 = op(e1, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f324, f1219])).
fof(f5494, plain, (spl24_116 | ~ spl24_127), inference(avatar_split_clause, [], [f5481, f1217, f1171])).
fof(f5481, plain, ((e0 = op(e0, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f322, f1219])).
fof(f5493, plain, (spl24_96 | ~ spl24_127), inference(avatar_split_clause, [], [f5480, f1217, f1087])).
fof(f5480, plain, ((e0 = op(e1, e0)) | ~ spl24_127), inference(backward_demodulation, [], [f321, f1219])).
fof(f5467, plain, (~ spl24_23 | ~ spl24_125 | spl24_202), inference(avatar_split_clause, [], [f5466, f1700, f1208, f780])).
fof(f5466, plain, (~ (e2 = op(e4, e0)) | (~ spl24_125 | spl24_202)), inference(forward_demodulation, [], [f1702, f1210])).
fof(f5447, plain, (~ spl24_89 | ~ spl24_61 | ~ spl24_72 | spl24_164), inference(avatar_split_clause, [], [f5341, f1494, f986, f940, f1057])).
fof(f5341, plain, (~ (e3 = op(e1, e2)) | (~ spl24_61 | ~ spl24_72 | spl24_164)), inference(backward_demodulation, [], [f5100, f988])).
fof(f5433, plain, (~ spl24_1 | ~ spl24_23 | ~ spl24_54 | spl24_161), inference(avatar_contradiction_clause, [], [f5432])).
fof(f5432, plain, ($false | (~ spl24_1 | ~ spl24_23 | ~ spl24_54 | spl24_161)), inference(subsumption_resolution, [], [f5431, f912])).
fof(f5431, plain, (~ (e3 = op(e2, e4)) | (~ spl24_1 | ~ spl24_23 | spl24_161)), inference(forward_demodulation, [], [f5424, f782])).
fof(f5424, plain, (~ (e3 = op(op(e4, e0), e4)) | (~ spl24_1 | spl24_161)), inference(backward_demodulation, [], [f1483, f690])).
fof(f5405, plain, (~ spl24_13 | ~ spl24_23), inference(avatar_split_clause, [], [f5400, f780, f738])).
fof(f5400, plain, (~ (e2 = op(e4, e2)) | ~ spl24_23), inference(backward_demodulation, [], [f473, f782])).
fof(f5371, plain, (~ spl24_31 | ~ spl24_46), inference(avatar_split_clause, [], [f5364, f877, f814])).
fof(f5364, plain, (~ (e0 = op(e3, e3)) | ~ spl24_46), inference(backward_demodulation, [], [f464, f879])).
fof(f5333, plain, (~ spl24_53 | ~ spl24_78), inference(avatar_split_clause, [], [f5328, f1011, f906])).
fof(f5328, plain, (~ (e2 = op(e2, e4)) | ~ spl24_78), inference(backward_demodulation, [], [f426, f1013])).
fof(f5309, plain, (~ spl24_31 | ~ spl24_106), inference(avatar_split_clause, [], [f5302, f1129, f814])).
fof(f5302, plain, (~ (e0 = op(e3, e3)) | ~ spl24_106), inference(backward_demodulation, [], [f414, f1131])).
fof(f5287, plain, (~ spl24_124 | ~ spl24_125), inference(avatar_contradiction_clause, [], [f5286])).
fof(f5286, plain, ($false | (~ spl24_124 | ~ spl24_125)), inference(subsumption_resolution, [], [f5285, f491])).
fof(f5285, plain, ((e3 = e4) | (~ spl24_124 | ~ spl24_125)), inference(backward_demodulation, [], [f1210, f1206])).
fof(f5238, plain, (~ spl24_100 | spl24_193 | ~ spl24_206), inference(avatar_contradiction_clause, [], [f5237])).
fof(f5237, plain, ($false | (~ spl24_100 | spl24_193 | ~ spl24_206)), inference(subsumption_resolution, [], [f5236, f1105])).
fof(f5236, plain, (~ (e4 = op(e1, e0)) | (spl24_193 | ~ spl24_206)), inference(forward_demodulation, [], [f1644, f1725])).
fof(f5235, plain, (~ spl24_118 | ~ spl24_91 | spl24_201), inference(avatar_split_clause, [], [f5234, f1690, f1066, f1179])).
fof(f5234, plain, (~ (e2 = op(e0, e1)) | (~ spl24_91 | spl24_201)), inference(forward_demodulation, [], [f1692, f1068])).
fof(f5226, plain, (~ spl24_119 | ~ spl24_91 | spl24_209), inference(avatar_split_clause, [], [f5225, f1746, f1066, f1183])).
fof(f5225, plain, (~ (e3 = op(e0, e1)) | (~ spl24_91 | spl24_209)), inference(forward_demodulation, [], [f1748, f1068])).
fof(f5207, plain, (~ spl24_40 | ~ spl24_61 | ~ spl24_74 | spl24_168), inference(avatar_contradiction_clause, [], [f5206])).
fof(f5206, plain, ($false | (~ spl24_40 | ~ spl24_61 | ~ spl24_74 | spl24_168)), inference(subsumption_resolution, [], [f5201, f853])).
fof(f5201, plain, (~ (e4 = op(e3, e2)) | (~ spl24_61 | ~ spl24_74 | spl24_168)), inference(backward_demodulation, [], [f5093, f996])).
fof(f5093, plain, (~ (e4 = op(op(e2, e0), e2)) | (~ spl24_61 | spl24_168)), inference(forward_demodulation, [], [f1514, f942])).
fof(f5156, plain, (~ spl24_72 | ~ spl24_123 | spl24_189), inference(avatar_split_clause, [], [f5155, f1624, f1200, f986])).
fof(f1624, plain, (spl24_189 <=> (e1 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_189])])).
fof(f5155, plain, (~ (e1 = op(e2, e0)) | (~ spl24_123 | spl24_189)), inference(forward_demodulation, [], [f1626, f1202])).
fof(f1626, plain, (~ (e1 = op(op(e0, e0), e0)) | spl24_189), inference(avatar_component_clause, [], [f1624])).
fof(f5136, plain, (~ spl24_113 | ~ spl24_13), inference(avatar_split_clause, [], [f5135, f738, f1158])).
fof(f5135, plain, (~ (e2 = op(e0, e2)) | ~ spl24_13), inference(forward_demodulation, [], [f405, f740])).
fof(f5133, plain, (~ spl24_114 | ~ spl24_119), inference(avatar_split_clause, [], [f5132, f1183, f1162])).
fof(f5132, plain, (~ (e3 = op(e0, e2)) | ~ spl24_119), inference(forward_demodulation, [], [f436, f1185])).
fof(f5125, plain, (~ spl24_108 | ~ spl24_31 | spl24_200), inference(avatar_split_clause, [], [f5042, f1684, f814, f1137])).
fof(f1684, plain, (spl24_200 <=> (e2 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl24_200])])).
fof(f5042, plain, (~ (e2 = op(e0, e3)) | (~ spl24_31 | spl24_200)), inference(backward_demodulation, [], [f1686, f816])).
fof(f1686, plain, (~ (e2 = op(op(e3, e3), e3)) | spl24_200), inference(avatar_component_clause, [], [f1684])).
fof(f5099, plain, (spl24_114 | ~ spl24_61 | ~ spl24_208), inference(avatar_split_clause, [], [f5098, f1740, f940, f1162])).
fof(f5098, plain, ((e3 = op(e0, e2)) | (~ spl24_61 | ~ spl24_208)), inference(forward_demodulation, [], [f1741, f942])).
fof(f5075, plain, (~ spl24_3 | ~ spl24_13), inference(avatar_split_clause, [], [f5072, f738, f696])).
fof(f5072, plain, (~ (e2 = op(e4, e4)) | ~ spl24_13), inference(backward_demodulation, [], [f480, f740])).
fof(f5025, plain, (~ spl24_3 | ~ spl24_53), inference(avatar_split_clause, [], [f5022, f906, f696])).
fof(f5022, plain, (~ (e2 = op(e4, e4)) | ~ spl24_53), inference(backward_demodulation, [], [f430, f908])).
fof(f5015, plain, (~ spl24_61 | ~ spl24_62), inference(avatar_contradiction_clause, [], [f5014])).
fof(f5014, plain, ($false | (~ spl24_61 | ~ spl24_62)), inference(subsumption_resolution, [], [f5013, f482])).
fof(f5013, plain, ((e0 = e1) | (~ spl24_61 | ~ spl24_62)), inference(backward_demodulation, [], [f946, f942])).
fof(f4987, plain, (~ spl24_119 | ~ spl24_120), inference(avatar_contradiction_clause, [], [f4986])).
fof(f4986, plain, ($false | (~ spl24_119 | ~ spl24_120)), inference(subsumption_resolution, [], [f4985, f491])).
fof(f4985, plain, ((e3 = e4) | (~ spl24_119 | ~ spl24_120)), inference(backward_demodulation, [], [f1189, f1185])).
fof(f4984, plain, (~ spl24_105 | ~ spl24_120), inference(avatar_split_clause, [], [f4983, f1187, f1124])).
fof(f4983, plain, (~ (e4 = op(e0, e4)) | ~ spl24_120), inference(backward_demodulation, [], [f438, f1189])).
fof(f4980, plain, (~ spl24_123 | ~ spl24_124), inference(avatar_contradiction_clause, [], [f4979])).
fof(f4979, plain, ($false | (~ spl24_123 | ~ spl24_124)), inference(subsumption_resolution, [], [f4978, f489])).
fof(f4978, plain, ((e2 = e3) | (~ spl24_123 | ~ spl24_124)), inference(backward_demodulation, [], [f1206, f1202])).
fof(f4948, plain, (~ spl24_62 | ~ spl24_70 | spl24_165), inference(avatar_contradiction_clause, [], [f4947])).
fof(f4947, plain, ($false | (~ spl24_62 | ~ spl24_70 | spl24_165)), inference(subsumption_resolution, [], [f4946, f979])).
fof(f4946, plain, (~ (e4 = op(e2, e1)) | (~ spl24_62 | spl24_165)), inference(forward_demodulation, [], [f1500, f946])).
fof(f4943, plain, (~ spl24_11 | ~ spl24_62 | ~ spl24_70 | spl24_195), inference(avatar_split_clause, [], [f4942, f1652, f977, f944, f730])).
fof(f4942, plain, (~ (e0 = op(e4, e2)) | (~ spl24_62 | ~ spl24_70 | spl24_195)), inference(forward_demodulation, [], [f4941, f979])).
fof(f4941, plain, (~ (e0 = op(op(e2, e1), e2)) | (~ spl24_62 | spl24_195)), inference(forward_demodulation, [], [f1654, f946])).
fof(f4939, plain, (~ spl24_74 | ~ spl24_124), inference(avatar_split_clause, [], [f4937, f1204, f994])).
fof(f4937, plain, (~ (e3 = op(e2, e0)) | ~ spl24_124), inference(backward_demodulation, [], [f383, f1206])).
fof(f4921, plain, (~ spl24_60 | ~ spl24_70), inference(avatar_split_clause, [], [f4920, f977, f935])).
fof(f4920, plain, (~ (e4 = op(e2, e3)) | ~ spl24_70), inference(forward_demodulation, [], [f457, f979])).
fof(f4911, plain, (~ spl24_47 | ~ spl24_124 | spl24_189), inference(avatar_split_clause, [], [f4786, f1624, f1204, f881])).
fof(f4786, plain, (~ (e1 = op(e3, e0)) | (~ spl24_124 | spl24_189)), inference(backward_demodulation, [], [f1626, f1206])).
fof(f4904, plain, (~ spl24_20 | ~ spl24_70), inference(avatar_split_clause, [], [f4903, f977, f767])).
fof(f4903, plain, (~ (e4 = op(e4, e1)) | ~ spl24_70), inference(forward_demodulation, [], [f400, f979])).
fof(f4901, plain, (~ spl24_19 | ~ spl24_91 | ~ spl24_100 | spl24_174), inference(avatar_split_clause, [], [f4827, f1540, f1103, f1066, f763])).
fof(f4827, plain, (~ (e3 = op(e4, e1)) | (~ spl24_91 | ~ spl24_100 | spl24_174)), inference(backward_demodulation, [], [f4605, f1105])).
fof(f4605, plain, (~ (e3 = op(op(e1, e0), e1)) | (~ spl24_91 | spl24_174)), inference(forward_demodulation, [], [f1542, f1068])).
fof(f4899, plain, (~ spl24_15 | ~ spl24_40), inference(avatar_split_clause, [], [f4898, f851, f746])).
fof(f4898, plain, (~ (e4 = op(e4, e2)) | ~ spl24_40), inference(forward_demodulation, [], [f411, f853])).
fof(f4852, plain, (~ spl24_73 | ~ spl24_74), inference(avatar_contradiction_clause, [], [f4851])).
fof(f4851, plain, ($false | (~ spl24_73 | ~ spl24_74)), inference(subsumption_resolution, [], [f4850, f489])).
fof(f4850, plain, ((e2 = e3) | (~ spl24_73 | ~ spl24_74)), inference(forward_demodulation, [], [f996, f992])).
fof(f4828, plain, (~ spl24_90 | ~ spl24_100), inference(avatar_split_clause, [], [f4822, f1103, f1061])).
fof(f4822, plain, (~ (e4 = op(e1, e2)) | ~ spl24_100), inference(backward_demodulation, [], [f443, f1105])).
fof(f4801, plain, (~ spl24_122 | ~ spl24_124), inference(avatar_contradiction_clause, [], [f4800])).
fof(f4800, plain, ($false | (~ spl24_122 | ~ spl24_124)), inference(subsumption_resolution, [], [f4799, f487])).
fof(f4799, plain, ((e1 = e3) | (~ spl24_122 | ~ spl24_124)), inference(backward_demodulation, [], [f1206, f1198])).
fof(f4797, plain, (~ spl24_108 | ~ spl24_124 | spl24_198), inference(avatar_split_clause, [], [f4790, f1668, f1204, f1137])).
fof(f4790, plain, (~ (e2 = op(e0, e3)) | (~ spl24_124 | spl24_198)), inference(backward_demodulation, [], [f1670, f1206])).
fof(f1670, plain, (~ (e2 = op(e0, op(e0, e0))) | spl24_198), inference(avatar_component_clause, [], [f1668])).
fof(f4780, plain, (~ spl24_7 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f4779])).
fof(f4779, plain, ($false | (~ spl24_7 | ~ spl24_129)), inference(subsumption_resolution, [], [f4778, f488])).
fof(f4778, plain, ((e1 = e4) | (~ spl24_7 | ~ spl24_129)), inference(forward_demodulation, [], [f4764, f715])).
fof(f4764, plain, ((e4 = op(e4, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f330, f1227])).
fof(f4732, plain, (~ spl24_120 | ~ spl24_91 | spl24_213), inference(avatar_split_clause, [], [f4505, f1786, f1066, f1187])).
fof(f4505, plain, (~ (e4 = op(e0, e1)) | (~ spl24_91 | spl24_213)), inference(forward_demodulation, [], [f1788, f1068])).
fof(f4723, plain, (~ spl24_99 | ~ spl24_91 | spl24_178), inference(avatar_split_clause, [], [f4592, f1558, f1066, f1099])).
fof(f4592, plain, (~ (e3 = op(e1, e0)) | (~ spl24_91 | spl24_178)), inference(forward_demodulation, [], [f1560, f1068])).
fof(f4722, plain, (~ spl24_98 | ~ spl24_91 | spl24_182), inference(avatar_split_clause, [], [f4599, f1580, f1066, f1095])).
fof(f4599, plain, (~ (e2 = op(e1, e0)) | (~ spl24_91 | spl24_182)), inference(forward_demodulation, [], [f1582, f1068])).
fof(f4720, plain, (~ spl24_90 | ~ spl24_40), inference(avatar_split_clause, [], [f4719, f851, f1061])).
fof(f4719, plain, (~ (e4 = op(e1, e2)) | ~ spl24_40), inference(forward_demodulation, [], [f407, f853])).
fof(f4715, plain, (~ spl24_85 | ~ spl24_60), inference(avatar_split_clause, [], [f4714, f935, f1040])).
fof(f4714, plain, (~ (e4 = op(e1, e3)) | ~ spl24_60), inference(forward_demodulation, [], [f416, f937])).
fof(f4687, plain, (~ spl24_35 | ~ spl24_40), inference(avatar_split_clause, [], [f4686, f851, f830])).
fof(f4686, plain, (~ (e4 = op(e3, e3)) | ~ spl24_40), inference(forward_demodulation, [], [f469, f853])).
fof(f4685, plain, (~ spl24_35 | ~ spl24_60), inference(avatar_split_clause, [], [f4684, f935, f830])).
fof(f4684, plain, (~ (e4 = op(e3, e3)) | ~ spl24_60), inference(forward_demodulation, [], [f419, f937])).
fof(f4682, plain, (~ spl24_3 | ~ spl24_11 | spl24_188), inference(avatar_contradiction_clause, [], [f4681])).
fof(f4681, plain, ($false | (~ spl24_3 | ~ spl24_11 | spl24_188)), inference(subsumption_resolution, [], [f4676, f732])).
fof(f4676, plain, (~ (e0 = op(e4, e2)) | (~ spl24_3 | spl24_188)), inference(backward_demodulation, [], [f1621, f698])).
fof(f1621, plain, (~ (e0 = op(e4, op(e4, e4))) | spl24_188), inference(avatar_component_clause, [], [f1619])).
fof(f4665, plain, (~ spl24_11 | ~ spl24_14), inference(avatar_contradiction_clause, [], [f4664])).
fof(f4664, plain, ($false | (~ spl24_11 | ~ spl24_14)), inference(subsumption_resolution, [], [f4663, f484])).
fof(f4663, plain, ((e0 = e3) | (~ spl24_11 | ~ spl24_14)), inference(forward_demodulation, [], [f744, f732])).
fof(f4651, plain, (~ spl24_43 | ~ spl24_45), inference(avatar_contradiction_clause, [], [f4650])).
fof(f4650, plain, ($false | (~ spl24_43 | ~ spl24_45)), inference(subsumption_resolution, [], [f4649, f490])).
fof(f4649, plain, ((e2 = e4) | (~ spl24_43 | ~ spl24_45)), inference(backward_demodulation, [], [f874, f866])).
fof(f4646, plain, (~ spl24_62 | spl24_86 | ~ spl24_163), inference(avatar_contradiction_clause, [], [f4645])).
fof(f4645, plain, ($false | (~ spl24_62 | spl24_86 | ~ spl24_163)), inference(subsumption_resolution, [], [f4639, f1046])).
fof(f1046, plain, (~ (e0 = op(e1, e2)) | spl24_86), inference(avatar_component_clause, [], [f1045])).
fof(f4639, plain, ((e0 = op(e1, e2)) | (~ spl24_62 | ~ spl24_163)), inference(backward_demodulation, [], [f1491, f946])).
fof(f1491, plain, ((e0 = op(op(e2, e2), e2)) | ~ spl24_163), inference(avatar_component_clause, [], [f1490])).
fof(f4644, plain, (~ spl24_57 | ~ spl24_62), inference(avatar_split_clause, [], [f4635, f944, f923])).
fof(f4635, plain, (~ (e1 = op(e2, e3)) | ~ spl24_62), inference(backward_demodulation, [], [f459, f946])).
fof(f4643, plain, (~ spl24_37 | ~ spl24_62), inference(avatar_split_clause, [], [f4634, f944, f839])).
fof(f839, plain, (spl24_37 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl24_37])])).
fof(f4634, plain, (~ (e1 = op(e3, e2)) | ~ spl24_62), inference(backward_demodulation, [], [f409, f946])).
fof(f4624, plain, (~ spl24_91 | ~ spl24_93), inference(avatar_contradiction_clause, [], [f4623])).
fof(f4623, plain, ($false | (~ spl24_91 | ~ spl24_93)), inference(subsumption_resolution, [], [f4622, f483])).
fof(f4622, plain, ((e0 = e2) | (~ spl24_91 | ~ spl24_93)), inference(forward_demodulation, [], [f1076, f1068])).
fof(f4621, plain, (~ spl24_108 | ~ spl24_109), inference(avatar_contradiction_clause, [], [f4620])).
fof(f4620, plain, ($false | (~ spl24_108 | ~ spl24_109)), inference(subsumption_resolution, [], [f4619, f489])).
fof(f4619, plain, ((e2 = e3) | (~ spl24_108 | ~ spl24_109)), inference(backward_demodulation, [], [f1143, f1139])).
fof(f4614, plain, (~ spl24_37 | ~ spl24_169 | spl24_179), inference(avatar_split_clause, [], [f4611, f1564, f1516, f839])).
fof(f4611, plain, (~ (e1 = op(e3, e2)) | (~ spl24_169 | spl24_179)), inference(backward_demodulation, [], [f1566, f1517])).
fof(f4613, plain, (~ spl24_40 | spl24_168 | ~ spl24_169), inference(avatar_split_clause, [], [f4610, f1516, f1512, f851])).
fof(f4610, plain, (~ (e4 = op(e3, e2)) | (spl24_168 | ~ spl24_169)), inference(backward_demodulation, [], [f1514, f1517])).
fof(f4582, plain, (~ spl24_90 | ~ spl24_85), inference(avatar_split_clause, [], [f4484, f1040, f1061])).
fof(f4484, plain, (~ (e4 = op(e1, e2)) | ~ spl24_85), inference(forward_demodulation, [], [f449, f1042])).
fof(f4580, plain, (~ spl24_61 | ~ spl24_11), inference(avatar_split_clause, [], [f4579, f730, f940])).
fof(f4579, plain, (~ (e0 = op(e2, e2)) | ~ spl24_11), inference(forward_demodulation, [], [f410, f732])).
fof(f4578, plain, (~ spl24_61 | ~ spl24_51), inference(avatar_split_clause, [], [f4577, f898, f940])).
fof(f898, plain, (spl24_51 <=> (e0 = op(e2, e4))), introduced(avatar_definition, [new_symbols(naming, [spl24_51])])).
fof(f4577, plain, (~ (e0 = op(e2, e2)) | ~ spl24_51), inference(forward_demodulation, [], [f460, f900])).
fof(f900, plain, ((e0 = op(e2, e4)) | ~ spl24_51), inference(avatar_component_clause, [], [f898])).
fof(f4567, plain, (~ spl24_59 | ~ spl24_109), inference(avatar_split_clause, [], [f4566, f1141, f931])).
fof(f4566, plain, (~ (e3 = op(e2, e3)) | ~ spl24_109), inference(forward_demodulation, [], [f413, f1143])).
fof(f4548, plain, (~ spl24_31 | ~ spl24_33), inference(avatar_contradiction_clause, [], [f4547])).
fof(f4547, plain, ($false | (~ spl24_31 | ~ spl24_33)), inference(subsumption_resolution, [], [f4546, f483])).
fof(f4546, plain, ((e0 = e2) | (~ spl24_31 | ~ spl24_33)), inference(forward_demodulation, [], [f824, f816])).
fof(f4545, plain, (~ spl24_61 | ~ spl24_65), inference(avatar_contradiction_clause, [], [f4544])).
fof(f4544, plain, ($false | (~ spl24_61 | ~ spl24_65)), inference(subsumption_resolution, [], [f4543, f485])).
fof(f4543, plain, ((e0 = e4) | (~ spl24_61 | ~ spl24_65)), inference(backward_demodulation, [], [f958, f942])).
fof(f4524, plain, (~ spl24_78 | ~ spl24_2 | spl24_199), inference(avatar_split_clause, [], [f4523, f1679, f692, f1011])).
fof(f4523, plain, (~ (e2 = op(e1, e4)) | (~ spl24_2 | spl24_199)), inference(forward_demodulation, [], [f1681, f694])).
fof(f4471, plain, (~ spl24_29 | ~ spl24_49), inference(avatar_split_clause, [], [f4470, f889, f805])).
fof(f4470, plain, (~ (e3 = op(e3, e4)) | ~ spl24_49), inference(forward_demodulation, [], [f465, f891])).
fof(f4464, plain, (~ spl24_6 | ~ spl24_31), inference(avatar_split_clause, [], [f4463, f814, f709])).
fof(f4463, plain, (~ (e0 = op(e4, e3)) | ~ spl24_31), inference(forward_demodulation, [], [f421, f816])).
fof(f4460, plain, (~ spl24_7 | ~ spl24_57), inference(avatar_split_clause, [], [f4459, f923, f713])).
fof(f4459, plain, (~ (e1 = op(e4, e3)) | ~ spl24_57), inference(forward_demodulation, [], [f420, f925])).
fof(f925, plain, ((e1 = op(e2, e3)) | ~ spl24_57), inference(avatar_component_clause, [], [f923])).
fof(f4458, plain, (~ spl24_8 | ~ spl24_18), inference(avatar_split_clause, [], [f4457, f759, f717])).
fof(f4457, plain, (~ (e2 = op(e4, e3)) | ~ spl24_18), inference(forward_demodulation, [], [f477, f761])).
fof(f4454, plain, (~ spl24_9 | ~ spl24_109), inference(avatar_split_clause, [], [f4453, f1141, f721])).
fof(f4453, plain, (~ (e3 = op(e4, e3)) | ~ spl24_109), inference(forward_demodulation, [], [f415, f1143])).
fof(f4452, plain, (~ spl24_6 | ~ spl24_11), inference(avatar_split_clause, [], [f4451, f730, f709])).
fof(f4451, plain, (~ (e0 = op(e4, e3)) | ~ spl24_11), inference(forward_demodulation, [], [f479, f732])).
fof(f4450, plain, (~ spl24_2 | ~ spl24_4), inference(avatar_contradiction_clause, [], [f4449])).
fof(f4449, plain, ($false | (~ spl24_2 | ~ spl24_4)), inference(subsumption_resolution, [], [f4448, f487])).
fof(f4448, plain, ((e1 = e3) | (~ spl24_2 | ~ spl24_4)), inference(forward_demodulation, [], [f702, f694])).
fof(f702, plain, ((e3 = op(e4, e4)) | ~ spl24_4), inference(avatar_component_clause, [], [f700])).
fof(f4443, plain, (~ spl24_31 | spl24_108 | ~ spl24_200), inference(avatar_contradiction_clause, [], [f4442])).
fof(f4442, plain, ($false | (~ spl24_31 | spl24_108 | ~ spl24_200)), inference(subsumption_resolution, [], [f4439, f1138])).
fof(f1138, plain, (~ (e2 = op(e0, e3)) | spl24_108), inference(avatar_component_clause, [], [f1137])).
fof(f4439, plain, ((e2 = op(e0, e3)) | (~ spl24_31 | ~ spl24_200)), inference(backward_demodulation, [], [f1685, f816])).
fof(f1685, plain, ((e2 = op(op(e3, e3), e3)) | ~ spl24_200), inference(avatar_component_clause, [], [f1684])).
fof(f4421, plain, (~ spl24_54 | ~ spl24_69), inference(avatar_split_clause, [], [f4420, f973, f910])).
fof(f4420, plain, (~ (e3 = op(e2, e4)) | ~ spl24_69), inference(backward_demodulation, [], [f458, f975])).
fof(f4410, plain, (~ spl24_83 | ~ spl24_85), inference(avatar_contradiction_clause, [], [f4409])).
fof(f4409, plain, ($false | (~ spl24_83 | ~ spl24_85)), inference(subsumption_resolution, [], [f4408, f490])).
fof(f4408, plain, ((e2 = e4) | (~ spl24_83 | ~ spl24_85)), inference(forward_demodulation, [], [f1042, f1034])).
fof(f4402, plain, (~ spl24_91 | ~ spl24_95), inference(avatar_contradiction_clause, [], [f4401])).
fof(f4401, plain, ($false | (~ spl24_91 | ~ spl24_95)), inference(subsumption_resolution, [], [f4400, f485])).
fof(f4400, plain, ((e0 = e4) | (~ spl24_91 | ~ spl24_95)), inference(backward_demodulation, [], [f1084, f1068])).
fof(f4385, plain, (spl24_51 | ~ spl24_65 | ~ spl24_197), inference(avatar_split_clause, [], [f4384, f1663, f956, f898])).
fof(f4384, plain, ((e0 = op(e2, e4)) | (~ spl24_65 | ~ spl24_197)), inference(forward_demodulation, [], [f1664, f958])).
fof(f4383, plain, (~ spl24_2 | spl24_78 | ~ spl24_199), inference(avatar_contradiction_clause, [], [f4382])).
fof(f4382, plain, ($false | (~ spl24_2 | spl24_78 | ~ spl24_199)), inference(subsumption_resolution, [], [f4381, f1012])).
fof(f1012, plain, (~ (e2 = op(e1, e4)) | spl24_78), inference(avatar_component_clause, [], [f1011])).
fof(f4381, plain, ((e2 = op(e1, e4)) | (~ spl24_2 | ~ spl24_199)), inference(forward_demodulation, [], [f1680, f694])).
fof(f4331, plain, (~ spl24_77 | ~ spl24_97), inference(avatar_split_clause, [], [f4330, f1091, f1007])).
fof(f4330, plain, (~ (e1 = op(e1, e4)) | ~ spl24_97), inference(forward_demodulation, [], [f445, f1093])).
fof(f4329, plain, (~ spl24_78 | ~ spl24_83), inference(avatar_split_clause, [], [f4328, f1032, f1011])).
fof(f4328, plain, (~ (e2 = op(e1, e4)) | ~ spl24_83), inference(forward_demodulation, [], [f451, f1034])).
fof(f4315, plain, (spl24_54 | ~ spl24_65 | ~ spl24_169), inference(avatar_split_clause, [], [f4151, f1516, f956, f910])).
fof(f4151, plain, ((e3 = op(e2, e4)) | (~ spl24_65 | ~ spl24_169)), inference(forward_demodulation, [], [f1517, f958])).
fof(f4313, plain, (~ spl24_32 | ~ spl24_37), inference(avatar_split_clause, [], [f4311, f839, f818])).
fof(f4311, plain, (~ (e1 = op(e3, e3)) | ~ spl24_37), inference(backward_demodulation, [], [f469, f841])).
fof(f841, plain, ((e1 = op(e3, e2)) | ~ spl24_37), inference(avatar_component_clause, [], [f839])).
fof(f4307, plain, (~ spl24_32 | ~ spl24_57), inference(avatar_split_clause, [], [f4306, f923, f818])).
fof(f4306, plain, (~ (e1 = op(e3, e3)) | ~ spl24_57), inference(forward_demodulation, [], [f419, f925])).
fof(f4296, plain, (~ spl24_15 | ~ spl24_25), inference(avatar_split_clause, [], [f4295, f788, f746])).
fof(f4295, plain, (~ (e4 = op(e4, e2)) | ~ spl24_25), inference(forward_demodulation, [], [f473, f790])).
fof(f4269, plain, (~ spl24_20 | ~ spl24_25), inference(avatar_split_clause, [], [f4267, f788, f767])).
fof(f4267, plain, (~ (e4 = op(e4, e1)) | ~ spl24_25), inference(backward_demodulation, [], [f472, f790])).
fof(f4239, plain, (~ spl24_76 | ~ spl24_79), inference(avatar_contradiction_clause, [], [f4238])).
fof(f4238, plain, ($false | (~ spl24_76 | ~ spl24_79)), inference(subsumption_resolution, [], [f4237, f484])).
fof(f4237, plain, ((e0 = e3) | (~ spl24_76 | ~ spl24_79)), inference(forward_demodulation, [], [f1017, f1005])).
fof(f4234, plain, (~ spl24_94 | ~ spl24_95), inference(avatar_contradiction_clause, [], [f4233])).
fof(f4233, plain, ($false | (~ spl24_94 | ~ spl24_95)), inference(subsumption_resolution, [], [f4232, f491])).
fof(f4232, plain, ((e3 = e4) | (~ spl24_94 | ~ spl24_95)), inference(forward_demodulation, [], [f1084, f1080])).
fof(f4231, plain, (~ spl24_87 | ~ spl24_97), inference(avatar_split_clause, [], [f4226, f1091, f1049])).
fof(f4226, plain, (~ (e1 = op(e1, e2)) | ~ spl24_97), inference(backward_demodulation, [], [f443, f1093])).
fof(f4230, plain, (~ spl24_47 | ~ spl24_97), inference(avatar_split_clause, [], [f4224, f1091, f881])).
fof(f4224, plain, (~ (e1 = op(e3, e0)) | ~ spl24_97), inference(backward_demodulation, [], [f387, f1093])).
fof(f4218, plain, (~ spl24_88 | ~ spl24_113), inference(avatar_split_clause, [], [f4213, f1158, f1053])).
fof(f4213, plain, (~ (e2 = op(e1, e2)) | ~ spl24_113), inference(backward_demodulation, [], [f402, f1160])).
fof(f4211, plain, (~ spl24_17 | ~ spl24_117), inference(avatar_split_clause, [], [f4204, f1175, f755])).
fof(f4204, plain, (~ (e1 = op(e4, e1)) | ~ spl24_117), inference(backward_demodulation, [], [f395, f1177])).
fof(f4210, plain, (~ spl24_67 | ~ spl24_117), inference(avatar_split_clause, [], [f4203, f1175, f965])).
fof(f4203, plain, (~ (e1 = op(e2, e1)) | ~ spl24_117), inference(backward_demodulation, [], [f393, f1177])).
fof(f4195, plain, (~ spl24_111 | ~ spl24_121), inference(avatar_split_clause, [], [f4185, f1192, f1150])).
fof(f4185, plain, (~ (e0 = op(e0, e2)) | ~ spl24_121), inference(backward_demodulation, [], [f433, f1194])).
fof(f4144, plain, (~ spl24_45 | ~ spl24_94 | spl24_213), inference(avatar_contradiction_clause, [], [f4143])).
fof(f4143, plain, ($false | (~ spl24_45 | ~ spl24_94 | spl24_213)), inference(subsumption_resolution, [], [f4142, f874])).
fof(f4142, plain, (~ (e4 = op(e3, e1)) | (~ spl24_94 | spl24_213)), inference(forward_demodulation, [], [f1788, f1080])).
fof(f4138, plain, (~ spl24_99 | spl24_190 | ~ spl24_206), inference(avatar_split_clause, [], [f4135, f1724, f1628, f1099])).
fof(f1628, plain, (spl24_190 <=> (e3 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl24_190])])).
fof(f4135, plain, (~ (e3 = op(e1, e0)) | (spl24_190 | ~ spl24_206)), inference(backward_demodulation, [], [f1630, f1725])).
fof(f1630, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | spl24_190), inference(avatar_component_clause, [], [f1628])).
fof(f4130, plain, (~ spl24_114 | ~ spl24_14), inference(avatar_split_clause, [], [f3928, f742, f1162])).
fof(f3928, plain, (~ (e3 = op(e0, e2)) | ~ spl24_14), inference(forward_demodulation, [], [f405, f744])).
fof(f4129, plain, (~ spl24_115 | ~ spl24_65), inference(avatar_split_clause, [], [f3931, f956, f1166])).
fof(f3931, plain, (~ (e4 = op(e0, e2)) | ~ spl24_65), inference(forward_demodulation, [], [f403, f958])).
fof(f4123, plain, (~ spl24_108 | ~ spl24_83), inference(avatar_split_clause, [], [f4122, f1032, f1137])).
fof(f4122, plain, (~ (e2 = op(e0, e3)) | ~ spl24_83), inference(forward_demodulation, [], [f412, f1034])).
fof(f4117, plain, (~ spl24_104 | ~ spl24_54), inference(avatar_split_clause, [], [f3885, f910, f1120])).
fof(f3885, plain, (~ (e3 = op(e0, e4)) | ~ spl24_54), inference(forward_demodulation, [], [f423, f912])).
fof(f4108, plain, (~ spl24_88 | ~ spl24_83), inference(avatar_split_clause, [], [f4107, f1032, f1053])).
fof(f4107, plain, (~ (e2 = op(e1, e2)) | ~ spl24_83), inference(forward_demodulation, [], [f449, f1034])).
fof(f4090, plain, (~ spl24_70 | spl24_177 | ~ spl24_182), inference(avatar_split_clause, [], [f3948, f1580, f1554, f977])).
fof(f3948, plain, (~ (e4 = op(e2, e1)) | (spl24_177 | ~ spl24_182)), inference(backward_demodulation, [], [f1556, f1581])).
fof(f4088, plain, (~ spl24_66 | ~ spl24_182 | spl24_203), inference(avatar_split_clause, [], [f3951, f1708, f1580, f961])).
fof(f3951, plain, (~ (e0 = op(e2, e1)) | (~ spl24_182 | spl24_203)), inference(backward_demodulation, [], [f1710, f1581])).
fof(f4085, plain, (~ spl24_50 | ~ spl24_45), inference(avatar_split_clause, [], [f4084, f872, f893])).
fof(f4084, plain, (~ (e4 = op(e3, e0)) | ~ spl24_45), inference(forward_demodulation, [], [f462, f874])).
fof(f4081, plain, (~ spl24_20 | ~ spl24_45), inference(avatar_split_clause, [], [f4080, f872, f767])).
fof(f4080, plain, (~ (e4 = op(e4, e1)) | ~ spl24_45), inference(backward_demodulation, [], [f401, f874])).
fof(f4077, plain, (~ spl24_41 | ~ spl24_32 | spl24_192), inference(avatar_split_clause, [], [f4051, f1637, f818, f856])).
fof(f4051, plain, (~ (e0 = op(e3, e1)) | (~ spl24_32 | spl24_192)), inference(backward_demodulation, [], [f1639, f820])).
fof(f4067, plain, (~ spl24_17 | ~ spl24_20), inference(avatar_contradiction_clause, [], [f4066])).
fof(f4066, plain, ($false | (~ spl24_17 | ~ spl24_20)), inference(subsumption_resolution, [], [f4064, f488])).
fof(f4064, plain, ((e1 = e4) | (~ spl24_17 | ~ spl24_20)), inference(backward_demodulation, [], [f769, f757])).
fof(f4061, plain, (~ spl24_32 | ~ spl24_83 | spl24_200), inference(avatar_contradiction_clause, [], [f4060])).
fof(f4060, plain, ($false | (~ spl24_32 | ~ spl24_83 | spl24_200)), inference(subsumption_resolution, [], [f4052, f1034])).
fof(f4052, plain, (~ (e2 = op(e1, e3)) | (~ spl24_32 | spl24_200)), inference(backward_demodulation, [], [f1686, f820])).
fof(f4041, plain, (~ spl24_41 | ~ spl24_45), inference(avatar_contradiction_clause, [], [f4040])).
fof(f4040, plain, ($false | (~ spl24_41 | ~ spl24_45)), inference(subsumption_resolution, [], [f4039, f485])).
fof(f4039, plain, ((e0 = e4) | (~ spl24_41 | ~ spl24_45)), inference(backward_demodulation, [], [f874, f858])).
fof(f4000, plain, (~ spl24_84 | ~ spl24_99), inference(avatar_split_clause, [], [f3997, f1099, f1036])).
fof(f3997, plain, (~ (e3 = op(e1, e3)) | ~ spl24_99), inference(backward_demodulation, [], [f444, f1101])).
fof(f3993, plain, (~ spl24_108 | ~ spl24_110), inference(avatar_contradiction_clause, [], [f3992])).
fof(f3992, plain, ($false | (~ spl24_108 | ~ spl24_110)), inference(subsumption_resolution, [], [f3991, f490])).
fof(f3991, plain, ((e2 = e4) | (~ spl24_108 | ~ spl24_110)), inference(forward_demodulation, [], [f1147, f1139])).
fof(f3982, plain, (~ spl24_14 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f3981])).
fof(f3981, plain, ($false | (~ spl24_14 | ~ spl24_128)), inference(subsumption_resolution, [], [f3980, f491])).
fof(f3980, plain, ((e3 = e4) | (~ spl24_14 | ~ spl24_128)), inference(forward_demodulation, [], [f3966, f744])).
fof(f3966, plain, ((e4 = op(e4, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f330, f1223])).
fof(f3979, plain, (~ spl24_54 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f3978])).
fof(f3978, plain, ($false | (~ spl24_54 | ~ spl24_128)), inference(subsumption_resolution, [], [f3977, f491])).
fof(f3977, plain, ((e3 = e4) | (~ spl24_54 | ~ spl24_128)), inference(forward_demodulation, [], [f3965, f912])).
fof(f3965, plain, ((e4 = op(e2, e4)) | ~ spl24_128), inference(backward_demodulation, [], [f329, f1223])).
fof(f3943, plain, (~ spl24_51 | ~ spl24_65 | spl24_197), inference(avatar_split_clause, [], [f3942, f1663, f956, f898])).
fof(f3942, plain, (~ (e0 = op(e2, e4)) | (~ spl24_65 | spl24_197)), inference(forward_demodulation, [], [f1665, f958])).
fof(f3917, plain, (~ spl24_98 | ~ spl24_23), inference(avatar_split_clause, [], [f3916, f780, f1095])).
fof(f3916, plain, (~ (e2 = op(e1, e0)) | ~ spl24_23), inference(forward_demodulation, [], [f388, f782])).
fof(f3914, plain, (~ spl24_98 | ~ spl24_102 | ~ spl24_125 | spl24_196), inference(avatar_split_clause, [], [f3816, f1658, f1208, f1112, f1095])).
fof(f3816, plain, (~ (e2 = op(e1, e0)) | (~ spl24_102 | ~ spl24_125 | spl24_196)), inference(backward_demodulation, [], [f3797, f1114])).
fof(f3797, plain, (~ (e2 = op(op(e0, e4), e0)) | (~ spl24_125 | spl24_196)), inference(forward_demodulation, [], [f1660, f1210])).
fof(f3913, plain, (~ spl24_99 | ~ spl24_102 | ~ spl24_125 | spl24_190), inference(avatar_split_clause, [], [f3817, f1628, f1208, f1112, f1099])).
fof(f3817, plain, (~ (e3 = op(e1, e0)) | (~ spl24_102 | ~ spl24_125 | spl24_190)), inference(backward_demodulation, [], [f3798, f1114])).
fof(f3798, plain, (~ (e3 = op(op(e0, e4), e0)) | (~ spl24_125 | spl24_190)), inference(forward_demodulation, [], [f1630, f1210])).
fof(f3908, plain, (~ spl24_85 | ~ spl24_80), inference(avatar_split_clause, [], [f3907, f1019, f1040])).
fof(f3907, plain, (~ (e4 = op(e1, e3)) | ~ spl24_80), inference(forward_demodulation, [], [f451, f1021])).
fof(f3897, plain, (~ spl24_63 | ~ spl24_68), inference(avatar_split_clause, [], [f3383, f969, f948])).
fof(f3383, plain, (~ (e2 = op(e2, e2)) | ~ spl24_68), inference(backward_demodulation, [], [f456, f971])).
fof(f3896, plain, (~ spl24_63 | ~ spl24_88), inference(avatar_split_clause, [], [f3468, f1053, f948])).
fof(f3468, plain, (~ (e2 = op(e2, e2)) | ~ spl24_88), inference(forward_demodulation, [], [f406, f1055])).
fof(f3890, plain, (~ spl24_36 | ~ spl24_54 | ~ spl24_65 | spl24_195), inference(avatar_split_clause, [], [f3889, f1652, f956, f910, f835])).
fof(f3889, plain, (~ (e0 = op(e3, e2)) | (~ spl24_54 | ~ spl24_65 | spl24_195)), inference(forward_demodulation, [], [f3686, f912])).
fof(f3686, plain, (~ (e0 = op(op(e2, e4), e2)) | (~ spl24_65 | spl24_195)), inference(backward_demodulation, [], [f1654, f958])).
fof(f3888, plain, (~ spl24_54 | ~ spl24_65 | spl24_169), inference(avatar_contradiction_clause, [], [f3887])).
fof(f3887, plain, ($false | (~ spl24_54 | ~ spl24_65 | spl24_169)), inference(subsumption_resolution, [], [f3682, f912])).
fof(f3682, plain, (~ (e3 = op(e2, e4)) | (~ spl24_65 | spl24_169)), inference(backward_demodulation, [], [f1518, f958])).
fof(f3872, plain, (~ spl24_39 | ~ spl24_14), inference(avatar_split_clause, [], [f3871, f742, f847])).
fof(f3871, plain, (~ (e3 = op(e3, e2)) | ~ spl24_14), inference(forward_demodulation, [], [f411, f744])).
fof(f3868, plain, (~ spl24_38 | ~ spl24_88), inference(avatar_split_clause, [], [f3867, f1053, f843])).
fof(f3867, plain, (~ (e2 = op(e3, e2)) | ~ spl24_88), inference(forward_demodulation, [], [f407, f1055])).
fof(f3865, plain, (spl24_54 | ~ spl24_3 | ~ spl24_207), inference(avatar_split_clause, [], [f3858, f1735, f696, f910])).
fof(f3858, plain, ((e3 = op(e2, e4)) | (~ spl24_3 | ~ spl24_207)), inference(backward_demodulation, [], [f1736, f698])).
fof(f1736, plain, ((e3 = op(op(e4, e4), e4)) | ~ spl24_207), inference(avatar_component_clause, [], [f1735])).
fof(f3862, plain, (~ spl24_52 | ~ spl24_3 | spl24_183), inference(avatar_split_clause, [], [f3854, f1591, f696, f902])).
fof(f3854, plain, (~ (e1 = op(e2, e4)) | (~ spl24_3 | spl24_183)), inference(backward_demodulation, [], [f1593, f698])).
fof(f3860, plain, (~ spl24_51 | ~ spl24_3 | spl24_155), inference(avatar_split_clause, [], [f3851, f1455, f696, f898])).
fof(f3851, plain, (~ (e0 = op(e2, e4)) | (~ spl24_3 | spl24_155)), inference(backward_demodulation, [], [f1457, f698])).
fof(f1457, plain, (~ (e0 = op(op(e4, e4), e4)) | spl24_155), inference(avatar_component_clause, [], [f1455])).
fof(f3826, plain, (~ spl24_51 | ~ spl24_54), inference(avatar_contradiction_clause, [], [f3825])).
fof(f3825, plain, ($false | (~ spl24_51 | ~ spl24_54)), inference(subsumption_resolution, [], [f3824, f484])).
fof(f3824, plain, ((e0 = e3) | (~ spl24_51 | ~ spl24_54)), inference(forward_demodulation, [], [f912, f900])).
fof(f3791, plain, (~ spl24_14 | ~ spl24_65 | spl24_208), inference(avatar_contradiction_clause, [], [f3790])).
fof(f3790, plain, ($false | (~ spl24_14 | ~ spl24_65 | spl24_208)), inference(subsumption_resolution, [], [f3789, f744])).
fof(f3789, plain, (~ (e3 = op(e4, e2)) | (~ spl24_65 | spl24_208)), inference(forward_demodulation, [], [f1742, f958])).
fof(f3781, plain, (~ spl24_105 | ~ spl24_125), inference(avatar_split_clause, [], [f3780, f1208, f1124])).
fof(f3780, plain, (~ (e4 = op(e0, e4)) | ~ spl24_125), inference(forward_demodulation, [], [f435, f1210])).
fof(f3735, plain, (~ spl24_2 | ~ spl24_3), inference(avatar_contradiction_clause, [], [f3734])).
fof(f3734, plain, ($false | (~ spl24_2 | ~ spl24_3)), inference(subsumption_resolution, [], [f3733, f486])).
fof(f3733, plain, ((e1 = e2) | (~ spl24_2 | ~ spl24_3)), inference(forward_demodulation, [], [f698, f694])).
fof(f3732, plain, (~ spl24_1 | ~ spl24_2), inference(avatar_contradiction_clause, [], [f3731])).
fof(f3731, plain, ($false | (~ spl24_1 | ~ spl24_2)), inference(subsumption_resolution, [], [f3730, f482])).
fof(f3730, plain, ((e0 = e1) | (~ spl24_1 | ~ spl24_2)), inference(forward_demodulation, [], [f694, f690])).
fof(f3729, plain, (~ spl24_1 | ~ spl24_3), inference(avatar_contradiction_clause, [], [f3728])).
fof(f3728, plain, ($false | (~ spl24_1 | ~ spl24_3)), inference(subsumption_resolution, [], [f3727, f483])).
fof(f3727, plain, ((e0 = e2) | (~ spl24_1 | ~ spl24_3)), inference(forward_demodulation, [], [f698, f690])).
fof(f3710, plain, (~ spl24_51 | ~ spl24_65 | ~ spl24_112 | spl24_179), inference(avatar_contradiction_clause, [], [f3709])).
fof(f3709, plain, ($false | (~ spl24_51 | ~ spl24_65 | ~ spl24_112 | spl24_179)), inference(subsumption_resolution, [], [f3703, f1156])).
fof(f3703, plain, (~ (e1 = op(e0, e2)) | (~ spl24_51 | ~ spl24_65 | spl24_179)), inference(backward_demodulation, [], [f3683, f900])).
fof(f3683, plain, (~ (e1 = op(op(e2, e4), e2)) | (~ spl24_65 | spl24_179)), inference(backward_demodulation, [], [f1566, f958])).
fof(f3696, plain, (~ spl24_52 | ~ spl24_57), inference(avatar_split_clause, [], [f3692, f923, f902])).
fof(f3692, plain, (~ (e1 = op(e2, e4)) | ~ spl24_57), inference(backward_demodulation, [], [f461, f925])).
fof(f3689, plain, (~ spl24_12 | ~ spl24_65 | spl24_185), inference(avatar_split_clause, [], [f3685, f1602, f956, f734])).
fof(f3685, plain, (~ (e1 = op(e4, e2)) | (~ spl24_65 | spl24_185)), inference(backward_demodulation, [], [f1604, f958])).
fof(f3688, plain, (~ spl24_52 | ~ spl24_65 | spl24_181), inference(avatar_split_clause, [], [f3684, f1575, f956, f902])).
fof(f3684, plain, (~ (e1 = op(e2, e4)) | (~ spl24_65 | spl24_181)), inference(backward_demodulation, [], [f1577, f958])).
fof(f3675, plain, (~ spl24_64 | ~ spl24_74), inference(avatar_split_clause, [], [f3670, f994, f952])).
fof(f3670, plain, (~ (e3 = op(e2, e2)) | ~ spl24_74), inference(backward_demodulation, [], [f453, f996])).
fof(f3668, plain, (~ spl24_8 | ~ spl24_108), inference(avatar_split_clause, [], [f3665, f1137, f717])).
fof(f3665, plain, (~ (e2 = op(e4, e3)) | ~ spl24_108), inference(backward_demodulation, [], [f415, f1139])).
fof(f3653, plain, (~ spl24_75 | ~ spl24_125), inference(avatar_split_clause, [], [f3644, f1208, f998])).
fof(f3644, plain, (~ (e4 = op(e2, e0)) | ~ spl24_125), inference(backward_demodulation, [], [f383, f1210])).
fof(f3641, plain, (~ spl24_57 | ~ spl24_28 | ~ spl24_35 | spl24_171), inference(avatar_split_clause, [], [f3640, f1526, f830, f801, f923])).
fof(f3640, plain, (~ (e1 = op(e2, e3)) | (~ spl24_28 | ~ spl24_35 | spl24_171)), inference(forward_demodulation, [], [f3639, f803])).
fof(f3639, plain, (~ (e1 = op(op(e3, e4), e3)) | (~ spl24_35 | spl24_171)), inference(forward_demodulation, [], [f1528, f832])).
fof(f3638, plain, (~ spl24_56 | ~ spl24_28 | ~ spl24_35 | spl24_187), inference(avatar_split_clause, [], [f3637, f1614, f830, f801, f919])).
fof(f3637, plain, (~ (e0 = op(e2, e3)) | (~ spl24_28 | ~ spl24_35 | spl24_187)), inference(forward_demodulation, [], [f3636, f803])).
fof(f3636, plain, (~ (e0 = op(op(e3, e4), e3)) | (~ spl24_35 | spl24_187)), inference(forward_demodulation, [], [f1616, f832])).
fof(f3635, plain, (~ spl24_8 | ~ spl24_35 | spl24_200), inference(avatar_split_clause, [], [f3634, f1684, f830, f717])).
fof(f3634, plain, (~ (e2 = op(e4, e3)) | (~ spl24_35 | spl24_200)), inference(forward_demodulation, [], [f1686, f832])).
fof(f3628, plain, (~ spl24_109 | ~ spl24_84), inference(avatar_split_clause, [], [f3627, f1036, f1141])).
fof(f3627, plain, (~ (e3 = op(e0, e3)) | ~ spl24_84), inference(forward_demodulation, [], [f412, f1038])).
fof(f1038, plain, ((e3 = op(e1, e3)) | ~ spl24_84), inference(avatar_component_clause, [], [f1036])).
fof(f3625, plain, (~ spl24_59 | ~ spl24_84), inference(avatar_split_clause, [], [f3624, f1036, f931])).
fof(f3624, plain, (~ (e3 = op(e2, e3)) | ~ spl24_84), inference(forward_demodulation, [], [f416, f1038])).
fof(f3621, plain, (~ spl24_51 | ~ spl24_1), inference(avatar_split_clause, [], [f3459, f688, f898])).
fof(f3459, plain, (~ (e0 = op(e2, e4)) | ~ spl24_1), inference(forward_demodulation, [], [f430, f690])).
fof(f3618, plain, (~ spl24_9 | ~ spl24_84), inference(avatar_split_clause, [], [f3617, f1036, f721])).
fof(f3617, plain, (~ (e3 = op(e4, e3)) | ~ spl24_84), inference(forward_demodulation, [], [f418, f1038])).
fof(f3613, plain, (~ spl24_33 | ~ spl24_35), inference(avatar_contradiction_clause, [], [f3612])).
fof(f3612, plain, ($false | (~ spl24_33 | ~ spl24_35)), inference(subsumption_resolution, [], [f3610, f490])).
fof(f3610, plain, ((e2 = e4) | (~ spl24_33 | ~ spl24_35)), inference(backward_demodulation, [], [f832, f824])).
fof(f3605, plain, (~ spl24_65 | ~ spl24_75), inference(avatar_split_clause, [], [f3604, f998, f956])).
fof(f3604, plain, (~ (e4 = op(e2, e2)) | ~ spl24_75), inference(backward_demodulation, [], [f453, f1000])).
fof(f3589, plain, (~ spl24_1 | ~ spl24_104 | spl24_207), inference(avatar_contradiction_clause, [], [f3588])).
fof(f3588, plain, ($false | (~ spl24_1 | ~ spl24_104 | spl24_207)), inference(subsumption_resolution, [], [f3587, f1122])).
fof(f3587, plain, (~ (e3 = op(e0, e4)) | (~ spl24_1 | spl24_207)), inference(forward_demodulation, [], [f1737, f690])).
fof(f3570, plain, (spl24_114 | ~ spl24_123 | ~ spl24_194), inference(avatar_split_clause, [], [f3569, f1646, f1200, f1162])).
fof(f3569, plain, ((e3 = op(e0, e2)) | (~ spl24_123 | ~ spl24_194)), inference(forward_demodulation, [], [f1647, f1202])).
fof(f3537, plain, (~ spl24_123 | ~ spl24_125), inference(avatar_contradiction_clause, [], [f3536])).
fof(f3536, plain, ($false | (~ spl24_123 | ~ spl24_125)), inference(subsumption_resolution, [], [f3535, f490])).
fof(f3535, plain, ((e2 = e4) | (~ spl24_123 | ~ spl24_125)), inference(backward_demodulation, [], [f1210, f1202])).
fof(f3487, plain, (~ spl24_114 | spl24_164 | ~ spl24_197), inference(avatar_split_clause, [], [f3271, f1663, f1494, f1162])).
fof(f3271, plain, (~ (e3 = op(e0, e2)) | (spl24_164 | ~ spl24_197)), inference(forward_demodulation, [], [f1496, f1664])).
fof(f3431, plain, (~ spl24_1 | ~ spl24_102 | ~ spl24_199), inference(avatar_contradiction_clause, [], [f3430])).
fof(f3430, plain, ($false | (~ spl24_1 | ~ spl24_102 | ~ spl24_199)), inference(subsumption_resolution, [], [f3429, f486])).
fof(f3429, plain, ((e1 = e2) | (~ spl24_1 | ~ spl24_102 | ~ spl24_199)), inference(forward_demodulation, [], [f3422, f1114])).
fof(f3422, plain, ((e2 = op(e0, e4)) | (~ spl24_1 | ~ spl24_199)), inference(backward_demodulation, [], [f1680, f690])).
fof(f3411, plain, (~ spl24_22 | ~ spl24_24), inference(avatar_contradiction_clause, [], [f3410])).
fof(f3410, plain, ($false | (~ spl24_22 | ~ spl24_24)), inference(subsumption_resolution, [], [f3409, f487])).
fof(f3409, plain, ((e1 = e3) | (~ spl24_22 | ~ spl24_24)), inference(forward_demodulation, [], [f786, f778])).
fof(f3402, plain, (~ spl24_37 | ~ spl24_47), inference(avatar_split_clause, [], [f3400, f881, f839])).
fof(f3400, plain, (~ (e1 = op(e3, e2)) | ~ spl24_47), inference(backward_demodulation, [], [f463, f883])).
fof(f3379, plain, (~ spl24_69 | ~ spl24_74), inference(avatar_split_clause, [], [f3376, f994, f973])).
fof(f3376, plain, (~ (e3 = op(e2, e1)) | ~ spl24_74), inference(backward_demodulation, [], [f452, f996])).
fof(f3374, plain, (~ spl24_4 | ~ spl24_79), inference(avatar_split_clause, [], [f3372, f1015, f700])).
fof(f3372, plain, (~ (e3 = op(e4, e4)) | ~ spl24_79), inference(backward_demodulation, [], [f428, f1017])).
fof(f428, plain, ~ (op(e1, e4) = op(e4, e4)), inference(cnf_transformation, [], [f4])).
fof(f3336, plain, (~ spl24_96 | ~ spl24_98), inference(avatar_contradiction_clause, [], [f3335])).
fof(f3335, plain, ($false | (~ spl24_96 | ~ spl24_98)), inference(subsumption_resolution, [], [f3333, f483])).
fof(f3333, plain, ((e0 = e2) | (~ spl24_96 | ~ spl24_98)), inference(backward_demodulation, [], [f1097, f1089])).
fof(f3332, plain, (~ spl24_98 | ~ spl24_99), inference(avatar_contradiction_clause, [], [f3331])).
fof(f3331, plain, ($false | (~ spl24_98 | ~ spl24_99)), inference(subsumption_resolution, [], [f3330, f489])).
fof(f3330, plain, ((e2 = e3) | (~ spl24_98 | ~ spl24_99)), inference(backward_demodulation, [], [f1101, f1097])).
fof(f3329, plain, (~ spl24_99 | ~ spl24_100), inference(avatar_contradiction_clause, [], [f3328])).
fof(f3328, plain, ($false | (~ spl24_99 | ~ spl24_100)), inference(subsumption_resolution, [], [f3327, f491])).
fof(f3327, plain, ((e3 = e4) | (~ spl24_99 | ~ spl24_100)), inference(forward_demodulation, [], [f1105, f1101])).
fof(f3326, plain, (~ spl24_102 | ~ spl24_104), inference(avatar_contradiction_clause, [], [f3325])).
fof(f3325, plain, ($false | (~ spl24_102 | ~ spl24_104)), inference(subsumption_resolution, [], [f3324, f487])).
fof(f3324, plain, ((e1 = e3) | (~ spl24_102 | ~ spl24_104)), inference(backward_demodulation, [], [f1122, f1114])).
fof(f3320, plain, (~ spl24_79 | ~ spl24_104), inference(avatar_split_clause, [], [f3316, f1120, f1015])).
fof(f3316, plain, (~ (e3 = op(e1, e4)) | ~ spl24_104), inference(backward_demodulation, [], [f422, f1122])).
fof(f3258, plain, (~ spl24_103 | ~ spl24_125 | spl24_198), inference(avatar_split_clause, [], [f3257, f1668, f1208, f1116])).
fof(f3257, plain, (~ (e2 = op(e0, e4)) | (~ spl24_125 | spl24_198)), inference(forward_demodulation, [], [f1670, f1210])).
fof(f3242, plain, (~ spl24_95 | ~ spl24_20), inference(avatar_split_clause, [], [f3010, f767, f1082])).
fof(f3010, plain, (~ (e4 = op(e1, e1)) | ~ spl24_20), inference(forward_demodulation, [], [f398, f769])).
fof(f3236, plain, (~ spl24_90 | ~ spl24_115), inference(avatar_split_clause, [], [f3235, f1166, f1061])).
fof(f3235, plain, (~ (e4 = op(e1, e2)) | ~ spl24_115), inference(forward_demodulation, [], [f402, f1168])).
fof(f3232, plain, (~ spl24_83 | ~ spl24_8), inference(avatar_split_clause, [], [f3094, f717, f1032])).
fof(f3094, plain, (~ (e2 = op(e1, e3)) | ~ spl24_8), inference(forward_demodulation, [], [f418, f719])).
fof(f719, plain, ((e2 = op(e4, e3)) | ~ spl24_8), inference(avatar_component_clause, [], [f717])).
fof(f3206, plain, (~ spl24_105 | ~ spl24_55), inference(avatar_split_clause, [], [f3205, f914, f1124])).
fof(f3205, plain, (~ (e4 = op(e0, e4)) | ~ spl24_55), inference(forward_demodulation, [], [f423, f916])).
fof(f3177, plain, (~ spl24_11 | ~ spl24_13), inference(avatar_contradiction_clause, [], [f3176])).
fof(f3176, plain, ($false | (~ spl24_11 | ~ spl24_13)), inference(subsumption_resolution, [], [f3175, f483])).
fof(f3175, plain, ((e0 = e2) | (~ spl24_11 | ~ spl24_13)), inference(forward_demodulation, [], [f740, f732])).
fof(f3154, plain, (~ spl24_121 | ~ spl24_125), inference(avatar_contradiction_clause, [], [f3153])).
fof(f3153, plain, ($false | (~ spl24_121 | ~ spl24_125)), inference(subsumption_resolution, [], [f3152, f485])).
fof(f3152, plain, ((e0 = e4) | (~ spl24_121 | ~ spl24_125)), inference(forward_demodulation, [], [f1210, f1194])).
fof(f3151, plain, (~ spl24_11 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f3150])).
fof(f3150, plain, ($false | (~ spl24_11 | ~ spl24_128)), inference(subsumption_resolution, [], [f3149, f485])).
fof(f3149, plain, ((e0 = e4) | (~ spl24_11 | ~ spl24_128)), inference(forward_demodulation, [], [f3125, f732])).
fof(f3125, plain, ((e4 = op(e4, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f330, f1223])).
fof(f3148, plain, (~ spl24_52 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f3147])).
fof(f3147, plain, ($false | (~ spl24_52 | ~ spl24_128)), inference(subsumption_resolution, [], [f3146, f488])).
fof(f3146, plain, ((e1 = e4) | (~ spl24_52 | ~ spl24_128)), inference(forward_demodulation, [], [f3124, f904])).
fof(f3124, plain, ((e4 = op(e2, e4)) | ~ spl24_128), inference(backward_demodulation, [], [f329, f1223])).
fof(f3145, plain, (~ spl24_37 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f3144])).
fof(f3144, plain, ($false | (~ spl24_37 | ~ spl24_128)), inference(subsumption_resolution, [], [f3143, f487])).
fof(f3143, plain, ((e1 = e3) | (~ spl24_37 | ~ spl24_128)), inference(forward_demodulation, [], [f3123, f841])).
fof(f3123, plain, ((e3 = op(e3, e2)) | ~ spl24_128), inference(backward_demodulation, [], [f328, f1223])).
fof(f3142, plain, (~ spl24_56 | ~ spl24_128), inference(avatar_contradiction_clause, [], [f3141])).
fof(f3141, plain, ($false | (~ spl24_56 | ~ spl24_128)), inference(subsumption_resolution, [], [f3140, f484])).
fof(f3140, plain, ((e0 = e3) | (~ spl24_56 | ~ spl24_128)), inference(forward_demodulation, [], [f3122, f921])).
fof(f3122, plain, ((e3 = op(e2, e3)) | ~ spl24_128), inference(backward_demodulation, [], [f327, f1223])).
fof(f3108, plain, (~ spl24_115 | ~ spl24_56 | ~ spl24_64 | spl24_168), inference(avatar_split_clause, [], [f2862, f1512, f952, f919, f1166])).
fof(f2862, plain, (~ (e4 = op(e0, e2)) | (~ spl24_56 | ~ spl24_64 | spl24_168)), inference(backward_demodulation, [], [f2856, f921])).
fof(f2856, plain, (~ (e4 = op(op(e2, e3), e2)) | (~ spl24_64 | spl24_168)), inference(forward_demodulation, [], [f1514, f954])).
fof(f3075, plain, (~ spl24_8 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f3074])).
fof(f3074, plain, ($false | (~ spl24_8 | ~ spl24_129)), inference(subsumption_resolution, [], [f3073, f490])).
fof(f3073, plain, ((e2 = e4) | (~ spl24_8 | ~ spl24_129)), inference(forward_demodulation, [], [f3050, f719])).
fof(f3050, plain, ((e4 = op(e4, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f330, f1227])).
fof(f3072, plain, (~ spl24_28 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f3071])).
fof(f3071, plain, ($false | (~ spl24_28 | ~ spl24_129)), inference(subsumption_resolution, [], [f3070, f490])).
fof(f3070, plain, ((e2 = e4) | (~ spl24_28 | ~ spl24_129)), inference(forward_demodulation, [], [f3049, f803])).
fof(f3049, plain, ((e4 = op(e3, e4)) | ~ spl24_129), inference(backward_demodulation, [], [f329, f1227])).
fof(f3065, plain, (~ spl24_56 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f3064])).
fof(f3064, plain, ($false | (~ spl24_56 | ~ spl24_129)), inference(subsumption_resolution, [], [f3063, f483])).
fof(f3063, plain, ((e0 = e2) | (~ spl24_56 | ~ spl24_129)), inference(forward_demodulation, [], [f3046, f921])).
fof(f3046, plain, ((e2 = op(e2, e3)) | ~ spl24_129), inference(backward_demodulation, [], [f326, f1227])).
fof(f3062, plain, (~ spl24_37 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f3061])).
fof(f3061, plain, ($false | (~ spl24_37 | ~ spl24_129)), inference(subsumption_resolution, [], [f3060, f486])).
fof(f3060, plain, ((e1 = e2) | (~ spl24_37 | ~ spl24_129)), inference(forward_demodulation, [], [f3045, f841])).
fof(f3045, plain, ((e2 = op(e3, e2)) | ~ spl24_129), inference(backward_demodulation, [], [f325, f1227])).
fof(f3058, plain, (~ spl24_41 | ~ spl24_129), inference(avatar_contradiction_clause, [], [f3057])).
fof(f3057, plain, ($false | (~ spl24_41 | ~ spl24_129)), inference(subsumption_resolution, [], [f3056, f482])).
fof(f3056, plain, ((e0 = e1) | (~ spl24_41 | ~ spl24_129)), inference(forward_demodulation, [], [f3043, f858])).
fof(f3043, plain, ((e1 = op(e3, e1)) | ~ spl24_129), inference(backward_demodulation, [], [f323, f1227])).
fof(f2989, plain, (~ spl24_70 | ~ spl24_75), inference(avatar_split_clause, [], [f2988, f998, f977])).
fof(f2988, plain, (~ (e4 = op(e2, e1)) | ~ spl24_75), inference(forward_demodulation, [], [f452, f1000])).
fof(f2977, plain, (~ spl24_68 | ~ spl24_93), inference(avatar_split_clause, [], [f2976, f1074, f969])).
fof(f2976, plain, (~ (e2 = op(e2, e1)) | ~ spl24_93), inference(forward_demodulation, [], [f396, f1076])).
fof(f2946, plain, (~ spl24_93 | ~ spl24_94), inference(avatar_contradiction_clause, [], [f2945])).
fof(f2945, plain, ($false | (~ spl24_93 | ~ spl24_94)), inference(subsumption_resolution, [], [f2944, f489])).
fof(f2944, plain, ((e2 = e3) | (~ spl24_93 | ~ spl24_94)), inference(backward_demodulation, [], [f1080, f1076])).
fof(f2943, plain, (~ spl24_107 | ~ spl24_109), inference(avatar_contradiction_clause, [], [f2942])).
fof(f2942, plain, ($false | (~ spl24_107 | ~ spl24_109)), inference(subsumption_resolution, [], [f2941, f487])).
fof(f2941, plain, ((e1 = e3) | (~ spl24_107 | ~ spl24_109)), inference(backward_demodulation, [], [f1143, f1135])).
fof(f2938, plain, (~ spl24_76 | ~ spl24_127), inference(avatar_contradiction_clause, [], [f2937])).
fof(f2937, plain, ($false | (~ spl24_76 | ~ spl24_127)), inference(subsumption_resolution, [], [f2936, f485])).
fof(f2936, plain, ((e0 = e4) | (~ spl24_76 | ~ spl24_127)), inference(forward_demodulation, [], [f2917, f1005])).
fof(f2917, plain, ((e4 = op(e1, e4)) | ~ spl24_127), inference(backward_demodulation, [], [f329, f1219])).
fof(f2935, plain, (~ spl24_41 | ~ spl24_127), inference(avatar_contradiction_clause, [], [f2934])).
fof(f2934, plain, ($false | (~ spl24_41 | ~ spl24_127)), inference(subsumption_resolution, [], [f2933, f484])).
fof(f2933, plain, ((e0 = e3) | (~ spl24_41 | ~ spl24_127)), inference(forward_demodulation, [], [f2916, f858])).
fof(f2916, plain, ((e3 = op(e3, e1)) | ~ spl24_127), inference(backward_demodulation, [], [f328, f1219])).
fof(f2894, plain, (spl24_8 | ~ spl24_35 | ~ spl24_200), inference(avatar_split_clause, [], [f2893, f1684, f830, f717])).
fof(f2893, plain, ((e2 = op(e4, e3)) | (~ spl24_35 | ~ spl24_200)), inference(forward_demodulation, [], [f1685, f832])).
fof(f2876, plain, (~ spl24_7 | ~ spl24_4 | spl24_172), inference(avatar_split_clause, [], [f2853, f1531, f700, f713])).
fof(f2853, plain, (~ (e1 = op(e4, e3)) | (~ spl24_4 | spl24_172)), inference(forward_demodulation, [], [f1533, f702])).
fof(f2829, plain, (~ spl24_3 | ~ spl24_4), inference(avatar_contradiction_clause, [], [f2828])).
fof(f2828, plain, ($false | (~ spl24_3 | ~ spl24_4)), inference(subsumption_resolution, [], [f2827, f489])).
fof(f2827, plain, ((e2 = e3) | (~ spl24_3 | ~ spl24_4)), inference(backward_demodulation, [], [f702, f698])).
fof(f2815, plain, (~ spl24_31 | ~ spl24_32), inference(avatar_contradiction_clause, [], [f2814])).
fof(f2814, plain, ($false | (~ spl24_31 | ~ spl24_32)), inference(subsumption_resolution, [], [f2813, f482])).
fof(f2813, plain, ((e0 = e1) | (~ spl24_31 | ~ spl24_32)), inference(backward_demodulation, [], [f820, f816])).
fof(f2804, plain, (~ spl24_51 | ~ spl24_52), inference(avatar_contradiction_clause, [], [f2803])).
fof(f2803, plain, ($false | (~ spl24_51 | ~ spl24_52)), inference(subsumption_resolution, [], [f2802, f482])).
fof(f2802, plain, ((e0 = e1) | (~ spl24_51 | ~ spl24_52)), inference(backward_demodulation, [], [f904, f900])).
fof(f2796, plain, (~ spl24_56 | ~ spl24_64 | spl24_197), inference(avatar_split_clause, [], [f2795, f1663, f952, f919])).
fof(f2795, plain, (~ (e0 = op(e2, e3)) | (~ spl24_64 | spl24_197)), inference(backward_demodulation, [], [f1665, f954])).
fof(f2771, plain, (~ spl24_93 | ~ spl24_83), inference(avatar_split_clause, [], [f2770, f1032, f1074])).
fof(f2770, plain, (~ (e2 = op(e1, e1)) | ~ spl24_83), inference(forward_demodulation, [], [f447, f1034])).
fof(f2758, plain, (~ spl24_43 | ~ spl24_28), inference(avatar_split_clause, [], [f2617, f801, f864])).
fof(f2617, plain, (~ (e2 = op(e3, e1)) | ~ spl24_28), inference(forward_demodulation, [], [f468, f803])).
fof(f2757, plain, (~ spl24_42 | ~ spl24_37), inference(avatar_split_clause, [], [f2756, f839, f860])).
fof(f2756, plain, (~ (e1 = op(e3, e1)) | ~ spl24_37), inference(forward_demodulation, [], [f466, f841])).
fof(f2755, plain, (~ spl24_32 | ~ spl24_35), inference(avatar_contradiction_clause, [], [f2754])).
fof(f2754, plain, ($false | (~ spl24_32 | ~ spl24_35)), inference(subsumption_resolution, [], [f2753, f488])).
fof(f2753, plain, ((e1 = e4) | (~ spl24_32 | ~ spl24_35)), inference(forward_demodulation, [], [f832, f820])).
fof(f2752, plain, (~ spl24_16 | ~ spl24_11), inference(avatar_split_clause, [], [f2751, f730, f751])).
fof(f2751, plain, (~ (e0 = op(e4, e1)) | ~ spl24_11), inference(forward_demodulation, [], [f476, f732])).
fof(f2734, plain, (~ spl24_32 | ~ spl24_34), inference(avatar_contradiction_clause, [], [f2733])).
fof(f2733, plain, ($false | (~ spl24_32 | ~ spl24_34)), inference(subsumption_resolution, [], [f2732, f487])).
fof(f2732, plain, ((e1 = e3) | (~ spl24_32 | ~ spl24_34)), inference(backward_demodulation, [], [f828, f820])).
fof(f2728, plain, (~ spl24_12 | ~ spl24_37), inference(avatar_split_clause, [], [f2727, f839, f734])).
fof(f2727, plain, (~ (e1 = op(e4, e2)) | ~ spl24_37), inference(backward_demodulation, [], [f411, f841])).
fof(f2713, plain, (~ spl24_62 | ~ spl24_69 | spl24_169), inference(avatar_contradiction_clause, [], [f2712])).
fof(f2712, plain, ($false | (~ spl24_62 | ~ spl24_69 | spl24_169)), inference(subsumption_resolution, [], [f2703, f975])).
fof(f2703, plain, (~ (e3 = op(e2, e1)) | (~ spl24_62 | spl24_169)), inference(backward_demodulation, [], [f1518, f946])).
fof(f2695, plain, (~ spl24_69 | ~ spl24_70), inference(avatar_contradiction_clause, [], [f2694])).
fof(f2694, plain, ($false | (~ spl24_69 | ~ spl24_70)), inference(subsumption_resolution, [], [f2693, f491])).
fof(f2693, plain, ((e3 = e4) | (~ spl24_69 | ~ spl24_70)), inference(backward_demodulation, [], [f979, f975])).
fof(f2683, plain, (~ spl24_89 | ~ spl24_90), inference(avatar_contradiction_clause, [], [f2682])).
fof(f2682, plain, ($false | (~ spl24_89 | ~ spl24_90)), inference(subsumption_resolution, [], [f2681, f491])).
fof(f2681, plain, ((e3 = e4) | (~ spl24_89 | ~ spl24_90)), inference(backward_demodulation, [], [f1063, f1059])).
fof(f2677, plain, (~ spl24_90 | ~ spl24_93 | spl24_175), inference(avatar_contradiction_clause, [], [f2676])).
fof(f2676, plain, ($false | (~ spl24_90 | ~ spl24_93 | spl24_175)), inference(subsumption_resolution, [], [f2670, f1063])).
fof(f2670, plain, (~ (e4 = op(e1, e2)) | (~ spl24_93 | spl24_175)), inference(backward_demodulation, [], [f1546, f1076])).
fof(f2673, plain, (~ spl24_18 | ~ spl24_93), inference(avatar_split_clause, [], [f2665, f1074, f759])).
fof(f2665, plain, (~ (e2 = op(e4, e1)) | ~ spl24_93), inference(backward_demodulation, [], [f398, f1076])).
fof(f2627, plain, (~ spl24_48 | ~ spl24_73), inference(avatar_split_clause, [], [f2626, f990, f885])).
fof(f2626, plain, (~ (e2 = op(e3, e0)) | ~ spl24_73), inference(forward_demodulation, [], [f389, f992])).
fof(f2622, plain, (~ spl24_50 | ~ spl24_25), inference(avatar_split_clause, [], [f2621, f788, f893])).
fof(f2621, plain, (~ (e4 = op(e3, e0)) | ~ spl24_25), inference(forward_demodulation, [], [f391, f790])).
fof(f2616, plain, (~ spl24_47 | ~ spl24_42), inference(avatar_split_clause, [], [f2615, f860, f881])).
fof(f2615, plain, (~ (e1 = op(e3, e0)) | ~ spl24_42), inference(forward_demodulation, [], [f462, f862])).
fof(f2607, plain, (~ spl24_38 | ~ spl24_28), inference(avatar_split_clause, [], [f2606, f801, f843])).
fof(f2606, plain, (~ (e2 = op(e3, e2)) | ~ spl24_28), inference(forward_demodulation, [], [f470, f803])).
fof(f2592, plain, (~ spl24_1 | ~ spl24_4), inference(avatar_contradiction_clause, [], [f2591])).
fof(f2591, plain, ($false | (~ spl24_1 | ~ spl24_4)), inference(subsumption_resolution, [], [f2590, f484])).
fof(f2590, plain, ((e0 = e3) | (~ spl24_1 | ~ spl24_4)), inference(backward_demodulation, [], [f702, f690])).
fof(f2583, plain, (~ spl24_36 | ~ spl24_38), inference(avatar_contradiction_clause, [], [f2582])).
fof(f2582, plain, ($false | (~ spl24_36 | ~ spl24_38)), inference(subsumption_resolution, [], [f2581, f483])).
fof(f2581, plain, ((e0 = e2) | (~ spl24_36 | ~ spl24_38)), inference(backward_demodulation, [], [f845, f837])).
fof(f2564, plain, (~ spl24_37 | ~ spl24_64 | spl24_185), inference(avatar_split_clause, [], [f2560, f1602, f952, f839])).
fof(f2560, plain, (~ (e1 = op(e3, e2)) | (~ spl24_64 | spl24_185)), inference(backward_demodulation, [], [f1604, f954])).
fof(f2563, plain, (~ spl24_57 | ~ spl24_64 | spl24_181), inference(avatar_contradiction_clause, [], [f2562])).
fof(f2562, plain, ($false | (~ spl24_57 | ~ spl24_64 | spl24_181)), inference(subsumption_resolution, [], [f2559, f925])).
fof(f2559, plain, (~ (e1 = op(e2, e3)) | (~ spl24_64 | spl24_181)), inference(backward_demodulation, [], [f1577, f954])).
fof(f2555, plain, (~ spl24_66 | ~ spl24_70), inference(avatar_contradiction_clause, [], [f2554])).
fof(f2554, plain, ($false | (~ spl24_66 | ~ spl24_70)), inference(subsumption_resolution, [], [f2553, f485])).
fof(f2553, plain, ((e0 = e4) | (~ spl24_66 | ~ spl24_70)), inference(backward_demodulation, [], [f979, f963])).
fof(f2547, plain, (~ spl24_53 | ~ spl24_73), inference(avatar_split_clause, [], [f2543, f990, f906])).
fof(f2543, plain, (~ (e2 = op(e2, e4)) | ~ spl24_73), inference(backward_demodulation, [], [f455, f992])).
fof(f2546, plain, (~ spl24_63 | ~ spl24_73), inference(avatar_split_clause, [], [f2542, f990, f948])).
fof(f2542, plain, (~ (e2 = op(e2, e2)) | ~ spl24_73), inference(backward_demodulation, [], [f453, f992])).
fof(f2523, plain, (~ spl24_63 | ~ spl24_113), inference(avatar_split_clause, [], [f2521, f1158, f948])).
fof(f2521, plain, (~ (e2 = op(e2, e2)) | ~ spl24_113), inference(backward_demodulation, [], [f403, f1160])).
fof(f2517, plain, (~ spl24_92 | ~ spl24_117), inference(avatar_split_clause, [], [f2514, f1175, f1070])).
fof(f2514, plain, (~ (e1 = op(e1, e1)) | ~ spl24_117), inference(backward_demodulation, [], [f392, f1177])).
fof(f2506, plain, (~ spl24_116 | ~ spl24_121), inference(avatar_split_clause, [], [f2497, f1192, f1171])).
fof(f2497, plain, (~ (e0 = op(e0, e1)) | ~ spl24_121), inference(backward_demodulation, [], [f432, f1194])).
fof(f2494, plain, (~ spl24_27 | ~ spl24_4 | spl24_183), inference(avatar_split_clause, [], [f2493, f1591, f700, f797])).
fof(f2493, plain, (~ (e1 = op(e3, e4)) | (~ spl24_4 | spl24_183)), inference(forward_demodulation, [], [f1593, f702])).
fof(f2492, plain, (~ spl24_28 | ~ spl24_4 | spl24_199), inference(avatar_split_clause, [], [f2491, f1679, f700, f801])).
fof(f2491, plain, (~ (e2 = op(e3, e4)) | (~ spl24_4 | spl24_199)), inference(forward_demodulation, [], [f1681, f702])).
fof(f2489, plain, (~ spl24_122 | ~ spl24_97), inference(avatar_split_clause, [], [f2488, f1091, f1196])).
fof(f2488, plain, (~ (op(e0, e0) = e1) | ~ spl24_97), inference(forward_demodulation, [], [f382, f1093])).
fof(f2485, plain, (~ spl24_119 | ~ spl24_44), inference(avatar_split_clause, [], [f2484, f868, f1183])).
fof(f2484, plain, (~ (e3 = op(e0, e1)) | ~ spl24_44), inference(forward_demodulation, [], [f394, f870])).
fof(f2460, plain, (~ spl24_69 | ~ spl24_44), inference(avatar_split_clause, [], [f2459, f868, f973])).
fof(f2459, plain, (~ (e3 = op(e2, e1)) | ~ spl24_44), inference(forward_demodulation, [], [f399, f870])).
fof(f2448, plain, (~ spl24_44 | ~ spl24_49), inference(avatar_contradiction_clause, [], [f2447])).
fof(f2447, plain, ($false | (~ spl24_44 | ~ spl24_49)), inference(subsumption_resolution, [], [f2446, f891])).
fof(f2446, plain, (~ (e3 = op(e3, e0)) | ~ spl24_44), inference(forward_demodulation, [], [f462, f870])).
fof(f2414, plain, (~ spl24_11 | ~ spl24_12), inference(avatar_contradiction_clause, [], [f2413])).
fof(f2413, plain, ($false | (~ spl24_11 | ~ spl24_12)), inference(subsumption_resolution, [], [f2412, f482])).
fof(f2412, plain, ((e0 = e1) | (~ spl24_11 | ~ spl24_12)), inference(backward_demodulation, [], [f736, f732])).
fof(f736, plain, ((e1 = op(e4, e2)) | ~ spl24_12), inference(avatar_component_clause, [], [f734])).
fof(f2399, plain, (~ spl24_39 | ~ spl24_40), inference(avatar_contradiction_clause, [], [f2398])).
fof(f2398, plain, ($false | (~ spl24_39 | ~ spl24_40)), inference(subsumption_resolution, [], [f2397, f491])).
fof(f2397, plain, ((e3 = e4) | (~ spl24_39 | ~ spl24_40)), inference(backward_demodulation, [], [f853, f849])).
fof(f2396, plain, (~ spl24_39 | ~ spl24_44), inference(avatar_split_clause, [], [f2395, f868, f847])).
fof(f2395, plain, (~ (e3 = op(e3, e2)) | ~ spl24_44), inference(backward_demodulation, [], [f466, f870])).
fof(f2388, plain, (~ spl24_51 | ~ spl24_53), inference(avatar_contradiction_clause, [], [f2387])).
fof(f2387, plain, ($false | (~ spl24_51 | ~ spl24_53)), inference(subsumption_resolution, [], [f2386, f483])).
fof(f2386, plain, ((e0 = e2) | (~ spl24_51 | ~ spl24_53)), inference(backward_demodulation, [], [f908, f900])).
fof(f2382, plain, (~ spl24_57 | ~ spl24_59), inference(avatar_contradiction_clause, [], [f2381])).
fof(f2381, plain, ($false | (~ spl24_57 | ~ spl24_59)), inference(subsumption_resolution, [], [f2380, f487])).
fof(f2380, plain, ((e1 = e3) | (~ spl24_57 | ~ spl24_59)), inference(backward_demodulation, [], [f933, f925])).
fof(f2379, plain, (~ spl24_59 | ~ spl24_60), inference(avatar_contradiction_clause, [], [f2378])).
fof(f2378, plain, ($false | (~ spl24_59 | ~ spl24_60)), inference(subsumption_resolution, [], [f2377, f491])).
fof(f2377, plain, ((e3 = e4) | (~ spl24_59 | ~ spl24_60)), inference(backward_demodulation, [], [f937, f933])).
fof(f2348, plain, (~ spl24_81 | ~ spl24_84), inference(avatar_contradiction_clause, [], [f2347])).
fof(f2347, plain, ($false | (~ spl24_81 | ~ spl24_84)), inference(subsumption_resolution, [], [f2346, f484])).
fof(f2346, plain, ((e0 = e3) | (~ spl24_81 | ~ spl24_84)), inference(forward_demodulation, [], [f1038, f1026])).
fof(f2335, plain, (~ spl24_87 | ~ spl24_92), inference(avatar_split_clause, [], [f2328, f1070, f1049])).
fof(f2328, plain, (~ (e1 = op(e1, e2)) | ~ spl24_92), inference(backward_demodulation, [], [f446, f1072])).
fof(f2324, plain, (~ spl24_92 | ~ spl24_97), inference(avatar_split_clause, [], [f2321, f1091, f1070])).
fof(f2321, plain, (~ (e1 = op(e1, e1)) | ~ spl24_97), inference(backward_demodulation, [], [f442, f1093])).
fof(f2323, plain, (~ spl24_72 | ~ spl24_97), inference(avatar_split_clause, [], [f2320, f1091, f986])).
fof(f2320, plain, (~ (e1 = op(e2, e0)) | ~ spl24_97), inference(backward_demodulation, [], [f386, f1093])).
fof(f2315, plain, (~ spl24_37 | ~ spl24_112), inference(avatar_split_clause, [], [f2312, f1154, f839])).
fof(f2312, plain, (~ (e1 = op(e3, e2)) | ~ spl24_112), inference(backward_demodulation, [], [f404, f1156])).
fof(f2301, plain, (~ spl24_107 | ~ spl24_122), inference(avatar_split_clause, [], [f2291, f1196, f1133])).
fof(f2291, plain, (~ (e1 = op(e0, e3)) | ~ spl24_122), inference(backward_demodulation, [], [f434, f1198])).
fof(f2287, plain, (~ spl24_23 | ~ spl24_126), inference(avatar_contradiction_clause, [], [f2286])).
fof(f2286, plain, ($false | (~ spl24_23 | ~ spl24_126)), inference(subsumption_resolution, [], [f2285, f490])).
fof(f2285, plain, ((e2 = e4) | (~ spl24_23 | ~ spl24_126)), inference(forward_demodulation, [], [f2268, f782])).
fof(f2268, plain, ((e4 = op(e4, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f330, f1215])).
fof(f2284, plain, (~ spl24_103 | ~ spl24_126), inference(avatar_contradiction_clause, [], [f2283])).
fof(f2283, plain, ($false | (~ spl24_103 | ~ spl24_126)), inference(subsumption_resolution, [], [f2282, f490])).
fof(f2282, plain, ((e2 = e4) | (~ spl24_103 | ~ spl24_126)), inference(forward_demodulation, [], [f2267, f1118])).
fof(f2267, plain, ((e4 = op(e0, e4)) | ~ spl24_126), inference(backward_demodulation, [], [f329, f1215])).
fof(f2281, plain, (~ spl24_46 | ~ spl24_126), inference(avatar_contradiction_clause, [], [f2280])).
fof(f2280, plain, ($false | (~ spl24_46 | ~ spl24_126)), inference(subsumption_resolution, [], [f2279, f484])).
fof(f2279, plain, ((e0 = e3) | (~ spl24_46 | ~ spl24_126)), inference(forward_demodulation, [], [f2266, f879])).
fof(f2266, plain, ((e3 = op(e3, e0)) | ~ spl24_126), inference(backward_demodulation, [], [f328, f1215])).
fof(f2274, plain, (~ spl24_116 | ~ spl24_126), inference(avatar_contradiction_clause, [], [f2273])).
fof(f2273, plain, ($false | (~ spl24_116 | ~ spl24_126)), inference(subsumption_resolution, [], [f2272, f482])).
fof(f2272, plain, ((e0 = e1) | (~ spl24_116 | ~ spl24_126)), inference(forward_demodulation, [], [f2261, f1173])).
fof(f2261, plain, ((e1 = op(e0, e1)) | ~ spl24_126), inference(backward_demodulation, [], [f323, f1215])).
fof(f2258, plain, (~ spl24_27 | ~ spl24_35 | ~ spl24_81 | spl24_187), inference(avatar_contradiction_clause, [], [f2257])).
fof(f2257, plain, ($false | (~ spl24_27 | ~ spl24_35 | ~ spl24_81 | spl24_187)), inference(subsumption_resolution, [], [f2256, f1026])).
fof(f2256, plain, (~ (e0 = op(e1, e3)) | (~ spl24_27 | ~ spl24_35 | spl24_187)), inference(forward_demodulation, [], [f2255, f799])).
fof(f2255, plain, (~ (e0 = op(op(e3, e4), e3)) | (~ spl24_35 | spl24_187)), inference(forward_demodulation, [], [f1616, f832])).
fof(f2193, plain, (~ spl24_68 | ~ spl24_53), inference(avatar_split_clause, [], [f2192, f906, f969])).
fof(f2192, plain, (~ (e2 = op(e2, e1)) | ~ spl24_53), inference(forward_demodulation, [], [f458, f908])).
fof(f2178, plain, (~ spl24_43 | ~ spl24_18), inference(avatar_split_clause, [], [f2177, f759, f864])).
fof(f2177, plain, (~ (e2 = op(e3, e1)) | ~ spl24_18), inference(forward_demodulation, [], [f401, f761])).
fof(f2168, plain, (~ spl24_37 | ~ spl24_27), inference(avatar_split_clause, [], [f2167, f797, f839])).
fof(f2167, plain, (~ (e1 = op(e3, e2)) | ~ spl24_27), inference(forward_demodulation, [], [f470, f799])).
fof(f2164, plain, (~ spl24_26 | ~ spl24_4 | spl24_155), inference(avatar_split_clause, [], [f2132, f1455, f700, f793])).
fof(f2132, plain, (~ (e0 = op(e3, e4)) | (~ spl24_4 | spl24_155)), inference(backward_demodulation, [], [f1457, f702])).
fof(f2153, plain, (~ spl24_18 | ~ spl24_20), inference(avatar_contradiction_clause, [], [f2152])).
fof(f2152, plain, ($false | (~ spl24_18 | ~ spl24_20)), inference(subsumption_resolution, [], [f2151, f490])).
fof(f2151, plain, ((e2 = e4) | (~ spl24_18 | ~ spl24_20)), inference(forward_demodulation, [], [f769, f761])).
fof(f2141, plain, (~ spl24_6 | ~ spl24_4 | spl24_188), inference(avatar_split_clause, [], [f2134, f1619, f700, f709])).
fof(f2134, plain, (~ (e0 = op(e4, e3)) | (~ spl24_4 | spl24_188)), inference(backward_demodulation, [], [f1621, f702])).
fof(f2128, plain, (~ spl24_7 | ~ spl24_8), inference(avatar_contradiction_clause, [], [f2127])).
fof(f2127, plain, ($false | (~ spl24_7 | ~ spl24_8)), inference(subsumption_resolution, [], [f2126, f486])).
fof(f2126, plain, ((e1 = e2) | (~ spl24_7 | ~ spl24_8)), inference(backward_demodulation, [], [f719, f715])).
fof(f2125, plain, (~ spl24_8 | ~ spl24_9), inference(avatar_contradiction_clause, [], [f2124])).
fof(f2124, plain, ($false | (~ spl24_8 | ~ spl24_9)), inference(subsumption_resolution, [], [f2123, f489])).
fof(f2123, plain, ((e2 = e3) | (~ spl24_8 | ~ spl24_9)), inference(backward_demodulation, [], [f723, f719])).
fof(f2113, plain, (~ spl24_4 | ~ spl24_14), inference(avatar_split_clause, [], [f2111, f742, f700])).
fof(f2111, plain, (~ (e3 = op(e4, e4)) | ~ spl24_14), inference(backward_demodulation, [], [f480, f744])).
fof(f2109, plain, (~ spl24_17 | ~ spl24_18), inference(avatar_contradiction_clause, [], [f2108])).
fof(f2108, plain, ($false | (~ spl24_17 | ~ spl24_18)), inference(subsumption_resolution, [], [f2107, f486])).
fof(f2107, plain, ((e1 = e2) | (~ spl24_17 | ~ spl24_18)), inference(backward_demodulation, [], [f761, f757])).
fof(f2106, plain, (~ spl24_18 | ~ spl24_19), inference(avatar_contradiction_clause, [], [f2105])).
fof(f2105, plain, ($false | (~ spl24_18 | ~ spl24_19)), inference(subsumption_resolution, [], [f2104, f489])).
fof(f2104, plain, ((e2 = e3) | (~ spl24_18 | ~ spl24_19)), inference(backward_demodulation, [], [f765, f761])).
fof(f2103, plain, (~ spl24_19 | ~ spl24_20), inference(avatar_contradiction_clause, [], [f2102])).
fof(f2102, plain, ($false | (~ spl24_19 | ~ spl24_20)), inference(subsumption_resolution, [], [f2101, f491])).
fof(f2101, plain, ((e3 = e4) | (~ spl24_19 | ~ spl24_20)), inference(backward_demodulation, [], [f769, f765])).
fof(f2094, plain, (~ spl24_22 | ~ spl24_23), inference(avatar_contradiction_clause, [], [f2093])).
fof(f2093, plain, ($false | (~ spl24_22 | ~ spl24_23)), inference(subsumption_resolution, [], [f2092, f486])).
fof(f2092, plain, ((e1 = e2) | (~ spl24_22 | ~ spl24_23)), inference(backward_demodulation, [], [f782, f778])).
fof(f2091, plain, (~ spl24_23 | ~ spl24_24), inference(avatar_contradiction_clause, [], [f2090])).
fof(f2090, plain, ($false | (~ spl24_23 | ~ spl24_24)), inference(subsumption_resolution, [], [f2089, f489])).
fof(f2089, plain, ((e2 = e3) | (~ spl24_23 | ~ spl24_24)), inference(backward_demodulation, [], [f786, f782])).
fof(f2088, plain, (~ spl24_24 | ~ spl24_25), inference(avatar_contradiction_clause, [], [f2087])).
fof(f2087, plain, ($false | (~ spl24_24 | ~ spl24_25)), inference(subsumption_resolution, [], [f2086, f491])).
fof(f2086, plain, ((e3 = e4) | (~ spl24_24 | ~ spl24_25)), inference(backward_demodulation, [], [f790, f786])).
fof(f2085, plain, (~ spl24_5 | ~ spl24_25), inference(avatar_split_clause, [], [f2081, f788, f704])).
fof(f2081, plain, (~ (e4 = op(e4, e4)) | ~ spl24_25), inference(backward_demodulation, [], [f475, f790])).
fof(f2077, plain, (~ spl24_26 | ~ spl24_27), inference(avatar_contradiction_clause, [], [f2076])).
fof(f2076, plain, ($false | (~ spl24_26 | ~ spl24_27)), inference(subsumption_resolution, [], [f2075, f482])).
fof(f2075, plain, ((e0 = e1) | (~ spl24_26 | ~ spl24_27)), inference(backward_demodulation, [], [f799, f795])).
fof(f2074, plain, (~ spl24_27 | ~ spl24_28), inference(avatar_contradiction_clause, [], [f2073])).
fof(f2073, plain, ($false | (~ spl24_27 | ~ spl24_28)), inference(subsumption_resolution, [], [f2072, f486])).
fof(f2072, plain, ((e1 = e2) | (~ spl24_27 | ~ spl24_28)), inference(backward_demodulation, [], [f803, f799])).
fof(f2071, plain, (~ spl24_28 | ~ spl24_29), inference(avatar_contradiction_clause, [], [f2070])).
fof(f2070, plain, ($false | (~ spl24_28 | ~ spl24_29)), inference(subsumption_resolution, [], [f2069, f489])).
fof(f2069, plain, ((e2 = e3) | (~ spl24_28 | ~ spl24_29)), inference(backward_demodulation, [], [f807, f803])).
fof(f2052, plain, (~ spl24_11 | ~ spl24_36), inference(avatar_split_clause, [], [f2049, f835, f730])).
fof(f2049, plain, (~ (e0 = op(e4, e2)) | ~ spl24_36), inference(backward_demodulation, [], [f411, f837])).
fof(f2047, plain, (~ spl24_31 | ~ spl24_41), inference(avatar_split_clause, [], [f2043, f856, f814])).
fof(f2043, plain, (~ (e0 = op(e3, e3)) | ~ spl24_41), inference(backward_demodulation, [], [f467, f858])).
fof(f2040, plain, (~ spl24_26 | ~ spl24_46), inference(avatar_split_clause, [], [f2035, f877, f793])).
fof(f2035, plain, (~ (e0 = op(e3, e4)) | ~ spl24_46), inference(backward_demodulation, [], [f465, f879])).
fof(f2024, plain, (~ spl24_37 | ~ spl24_54 | ~ spl24_65 | spl24_179), inference(avatar_split_clause, [], [f2021, f1564, f956, f910, f839])).
fof(f2021, plain, (~ (e1 = op(e3, e2)) | (~ spl24_54 | ~ spl24_65 | spl24_179)), inference(backward_demodulation, [], [f2004, f912])).
fof(f2004, plain, (~ (e1 = op(op(e2, e4), e2)) | (~ spl24_65 | spl24_179)), inference(backward_demodulation, [], [f1566, f958])).
fof(f2023, plain, (~ spl24_4 | ~ spl24_54), inference(avatar_split_clause, [], [f2020, f910, f700])).
fof(f2020, plain, (~ (e3 = op(e4, e4)) | ~ spl24_54), inference(backward_demodulation, [], [f430, f912])).
fof(f2018, plain, (~ spl24_51 | ~ spl24_56), inference(avatar_split_clause, [], [f2015, f919, f898])).
fof(f2015, plain, (~ (e0 = op(e2, e4)) | ~ spl24_56), inference(backward_demodulation, [], [f461, f921])).
fof(f2017, plain, (~ spl24_6 | ~ spl24_56), inference(avatar_split_clause, [], [f2014, f919, f709])).
fof(f2014, plain, (~ (e0 = op(e4, e3)) | ~ spl24_56), inference(backward_demodulation, [], [f420, f921])).
fof(f2009, plain, (~ spl24_60 | ~ spl24_65), inference(avatar_split_clause, [], [f2002, f956, f935])).
fof(f2002, plain, (~ (e4 = op(e2, e3)) | ~ spl24_65), inference(backward_demodulation, [], [f459, f958])).
fof(f1999, plain, (~ spl24_51 | ~ spl24_66), inference(avatar_split_clause, [], [f1994, f961, f898])).
fof(f1994, plain, (~ (e0 = op(e2, e4)) | ~ spl24_66), inference(backward_demodulation, [], [f458, f963])).
fof(f1997, plain, (~ spl24_61 | ~ spl24_66), inference(avatar_split_clause, [], [f1992, f961, f940])).
fof(f1992, plain, (~ (e0 = op(e2, e2)) | ~ spl24_66), inference(backward_demodulation, [], [f456, f963])).
fof(f1995, plain, (~ spl24_41 | ~ spl24_66), inference(avatar_split_clause, [], [f1990, f961, f856])).
fof(f1990, plain, (~ (e0 = op(e3, e1)) | ~ spl24_66), inference(backward_demodulation, [], [f399, f963])).
fof(f1989, plain, (~ spl24_51 | ~ spl24_71), inference(avatar_split_clause, [], [f1983, f982, f898])).
fof(f1983, plain, (~ (e0 = op(e2, e4)) | ~ spl24_71), inference(backward_demodulation, [], [f455, f984])).
fof(f1986, plain, (~ spl24_66 | ~ spl24_71), inference(avatar_split_clause, [], [f1980, f982, f961])).
fof(f1980, plain, (~ (e0 = op(e2, e1)) | ~ spl24_71), inference(backward_demodulation, [], [f452, f984])).
fof(f1985, plain, (~ spl24_21 | ~ spl24_71), inference(avatar_split_clause, [], [f1979, f982, f772])).
fof(f1979, plain, (~ (e0 = op(e4, e0)) | ~ spl24_71), inference(backward_demodulation, [], [f390, f984])).
fof(f1974, plain, (~ spl24_78 | ~ spl24_79), inference(avatar_contradiction_clause, [], [f1973])).
fof(f1973, plain, ($false | (~ spl24_78 | ~ spl24_79)), inference(subsumption_resolution, [], [f1972, f489])).
fof(f1972, plain, ((e2 = e3) | (~ spl24_78 | ~ spl24_79)), inference(backward_demodulation, [], [f1017, f1013])).
fof(f1971, plain, (~ spl24_79 | ~ spl24_80), inference(avatar_contradiction_clause, [], [f1970])).
fof(f1970, plain, ($false | (~ spl24_79 | ~ spl24_80)), inference(subsumption_resolution, [], [f1969, f491])).
fof(f1969, plain, ((e3 = e4) | (~ spl24_79 | ~ spl24_80)), inference(backward_demodulation, [], [f1021, f1017])).
fof(f1951, plain, (~ spl24_36 | ~ spl24_86), inference(avatar_split_clause, [], [f1946, f1045, f835])).
fof(f1946, plain, (~ (e0 = op(e3, e2)) | ~ spl24_86), inference(backward_demodulation, [], [f407, f1047])).
fof(f1940, plain, (~ spl24_76 | ~ spl24_91), inference(avatar_split_clause, [], [f1930, f1066, f1003])).
fof(f1930, plain, (~ (e0 = op(e1, e4)) | ~ spl24_91), inference(backward_demodulation, [], [f448, f1068])).
fof(f1935, plain, (~ spl24_66 | ~ spl24_91), inference(avatar_split_clause, [], [f1925, f1066, f961])).
fof(f1925, plain, (~ (e0 = op(e2, e1)) | ~ spl24_91), inference(backward_demodulation, [], [f396, f1068])).
fof(f1922, plain, (~ spl24_86 | ~ spl24_96), inference(avatar_split_clause, [], [f1915, f1087, f1045])).
fof(f1915, plain, (~ (e0 = op(e1, e2)) | ~ spl24_96), inference(backward_demodulation, [], [f443, f1089])).
fof(f1920, plain, (~ spl24_21 | ~ spl24_96), inference(avatar_split_clause, [], [f1913, f1087, f772])).
fof(f1913, plain, (~ (e0 = op(e4, e0)) | ~ spl24_96), inference(backward_demodulation, [], [f388, f1089])).
fof(f1918, plain, (~ spl24_71 | ~ spl24_96), inference(avatar_split_clause, [], [f1911, f1087, f982])).
fof(f1911, plain, (~ (e0 = op(e2, e0)) | ~ spl24_96), inference(backward_demodulation, [], [f386, f1089])).
fof(f1910, plain, (~ spl24_102 | ~ spl24_103), inference(avatar_contradiction_clause, [], [f1909])).
fof(f1909, plain, ($false | (~ spl24_102 | ~ spl24_103)), inference(subsumption_resolution, [], [f1908, f486])).
fof(f1908, plain, ((e1 = e2) | (~ spl24_102 | ~ spl24_103)), inference(backward_demodulation, [], [f1118, f1114])).
fof(f1907, plain, (~ spl24_103 | ~ spl24_104), inference(avatar_contradiction_clause, [], [f1906])).
fof(f1906, plain, ($false | (~ spl24_103 | ~ spl24_104)), inference(subsumption_resolution, [], [f1905, f489])).
fof(f1905, plain, ((e2 = e3) | (~ spl24_103 | ~ spl24_104)), inference(backward_demodulation, [], [f1122, f1118])).
fof(f1904, plain, (~ spl24_104 | ~ spl24_105), inference(avatar_contradiction_clause, [], [f1903])).
fof(f1903, plain, ($false | (~ spl24_104 | ~ spl24_105)), inference(subsumption_resolution, [], [f1902, f491])).
fof(f1902, plain, ((e3 = e4) | (~ spl24_104 | ~ spl24_105)), inference(backward_demodulation, [], [f1126, f1122])).
fof(f1898, plain, (~ spl24_80 | ~ spl24_105), inference(avatar_split_clause, [], [f1894, f1124, f1019])).
fof(f1894, plain, (~ (e4 = op(e1, e4)) | ~ spl24_105), inference(backward_demodulation, [], [f422, f1126])).
fof(f1892, plain, (~ spl24_6 | ~ spl24_106), inference(avatar_split_clause, [], [f1887, f1129, f709])).
fof(f1887, plain, (~ (e0 = op(e4, e3)) | ~ spl24_106), inference(backward_demodulation, [], [f415, f1131])).
fof(f1880, plain, (~ spl24_36 | ~ spl24_111), inference(avatar_split_clause, [], [f1874, f1150, f835])).
fof(f1874, plain, (~ (e0 = op(e3, e2)) | ~ spl24_111), inference(backward_demodulation, [], [f404, f1152])).
fof(f1871, plain, (~ spl24_101 | ~ spl24_116), inference(avatar_split_clause, [], [f1864, f1171, f1108])).
fof(f1864, plain, (~ (e0 = op(e0, e4)) | ~ spl24_116), inference(backward_demodulation, [], [f438, f1173])).
fof(f1869, plain, (~ spl24_111 | ~ spl24_116), inference(avatar_split_clause, [], [f1862, f1171, f1150])).
fof(f1862, plain, (~ (e0 = op(e0, e2)) | ~ spl24_116), inference(backward_demodulation, [], [f436, f1173])).
fof(f1868, plain, (~ spl24_16 | ~ spl24_116), inference(avatar_split_clause, [], [f1861, f1171, f751])).
fof(f1861, plain, (~ (e0 = op(e4, e1)) | ~ spl24_116), inference(backward_demodulation, [], [f395, f1173])).
fof(f1866, plain, (~ spl24_66 | ~ spl24_116), inference(avatar_split_clause, [], [f1859, f1171, f961])).
fof(f1859, plain, (~ (e0 = op(e2, e1)) | ~ spl24_116), inference(backward_demodulation, [], [f393, f1173])).
fof(f1853, plain, (~ spl24_101 | ~ spl24_121), inference(avatar_split_clause, [], [f1841, f1192, f1108])).
fof(f1841, plain, (~ (e0 = op(e0, e4)) | ~ spl24_121), inference(backward_demodulation, [], [f435, f1194])).
fof(f1849, plain, (~ spl24_21 | ~ spl24_121), inference(avatar_split_clause, [], [f1837, f1192, f772])).
fof(f1837, plain, (~ (e0 = op(e4, e0)) | ~ spl24_121), inference(backward_demodulation, [], [f385, f1194])).
fof(f1813, plain, (~ spl24_214 | ~ spl24_124 | ~ spl24_196 | ~ spl24_206), inference(avatar_split_clause, [], [f686, f1724, f1658, f1204, f1796])).
fof(f686, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e4 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f246])).
fof(f246, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e4 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f126])).
fof(f126, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e4 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax126)).
fof(f1812, plain, (~ spl24_213 | ~ spl24_94 | ~ spl24_180 | ~ spl24_205), inference(avatar_split_clause, [], [f685, f1719, f1570, f1078, f1786])).
fof(f685, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f245])).
fof(f245, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f125])).
fof(f125, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1)) & (e4 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax125)).
fof(f1811, plain, (~ spl24_214 | ~ spl24_124 | ~ spl24_204 | ~ spl24_198), inference(avatar_split_clause, [], [f684, f1668, f1714, f1204, f1796])).
fof(f684, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e4 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f244])).
fof(f244, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e4 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f124])).
fof(f124, plain, ~ ((e2 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e4 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax124)).
fof(f1810, plain, (~ spl24_212 | ~ spl24_64 | ~ spl24_179 | ~ spl24_197), inference(avatar_split_clause, [], [f683, f1663, f1564, f952, f1780])).
fof(f683, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f243])).
fof(f243, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f123])).
fof(f123, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, e2)) & (e4 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax123)).
fof(f1809, plain, (~ spl24_213 | ~ spl24_94 | ~ spl24_203 | ~ spl24_182), inference(avatar_split_clause, [], [f682, f1580, f1708, f1078, f1786])).
fof(f682, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f242])).
fof(f242, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f122])).
fof(f122, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1)) & (e4 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax122)).
fof(f1808, plain, (~ spl24_212 | ~ spl24_64 | ~ spl24_195 | ~ spl24_181), inference(avatar_split_clause, [], [f681, f1575, f1652, f952, f1780])).
fof(f681, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f241])).
fof(f241, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f121])).
fof(f121, plain, ~ ((e1 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, e2)) & (e4 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax121)).
fof(f1807, plain, (~ spl24_214 | ~ spl24_123 | ~ spl24_190 | ~ spl24_206), inference(avatar_split_clause, [], [f680, f1724, f1628, f1200, f1796])).
fof(f680, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e4 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f240])).
fof(f240, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e4 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f120])).
fof(f120, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e3 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e2) & (e4 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax120)).
fof(f1806, plain, (~ spl24_213 | ~ spl24_93 | ~ spl24_174 | ~ spl24_205), inference(avatar_split_clause, [], [f679, f1719, f1540, f1074, f1786])).
fof(f679, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f239])).
fof(f239, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f119])).
fof(f119, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1)) & (e4 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax119)).
fof(f1804, plain, (~ spl24_211 | ~ spl24_33 | ~ spl24_171 | ~ spl24_192), inference(avatar_split_clause, [], [f677, f1637, f1526, f822, f1775])).
fof(f677, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e4 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f237])).
fof(f237, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e2 = op(e3, e3)) | ~ (e4 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f117])).
fof(f117, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), e3)) & (e2 = op(e3, e3)) & (e4 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax117)).
fof(f1800, plain, (~ spl24_212 | ~ spl24_62 | ~ spl24_164 | ~ spl24_197), inference(avatar_split_clause, [], [f673, f1663, f1494, f944, f1780])).
fof(f673, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f233])).
fof(f233, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f113])).
fof(f113, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e3 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e4 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax113)).
fof(f1792, plain, (~ spl24_211 | ~ spl24_32 | ~ spl24_187 | ~ spl24_167), inference(avatar_split_clause, [], [f669, f1507, f1614, f818, f1775])).
fof(f669, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e4 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f229])).
fof(f229, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e4 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f109])).
fof(f109, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), e3)) & (e1 = op(e3, e3)) & (e4 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax109)).
fof(f1791, plain, (~ spl24_213 | ~ spl24_91 | ~ spl24_174 | ~ spl24_182), inference(avatar_split_clause, [], [f668, f1580, f1540, f1066, f1786])).
fof(f668, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f228])).
fof(f228, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f108])).
fof(f108, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e4 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax108)).
fof(f1790, plain, (~ spl24_212 | ~ spl24_61 | ~ spl24_164 | ~ spl24_181), inference(avatar_split_clause, [], [f667, f1575, f1494, f940, f1780])).
fof(f667, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f227])).
fof(f227, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f107])).
fof(f107, plain, ~ ((e1 = op(e2, op(e2, e2))) & (e3 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, e2)) & (e4 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax107)).
fof(f1789, plain, (~ spl24_213 | ~ spl24_91 | ~ spl24_180 | ~ spl24_178), inference(avatar_split_clause, [], [f666, f1558, f1570, f1066, f1786])).
fof(f666, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f226])).
fof(f226, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e4 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f106])).
fof(f106, plain, ~ ((e3 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e4 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax106)).
fof(f1783, plain, (~ spl24_212 | ~ spl24_61 | ~ spl24_179 | ~ spl24_169), inference(avatar_split_clause, [], [f664, f1516, f1564, f940, f1780])).
fof(f664, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f224])).
fof(f224, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e4 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f104])).
fof(f104, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, e2)) & (e4 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax104)).
fof(f1773, plain, (~ spl24_210 | ~ spl24_125 | ~ spl24_196 | ~ spl24_206), inference(avatar_split_clause, [], [f662, f1724, f1658, f1208, f1756])).
fof(f662, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e3 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f222])).
fof(f222, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e3 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f102])).
fof(f102, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e4) & (e3 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax102)).
fof(f1772, plain, (~ spl24_209 | ~ spl24_95 | ~ spl24_180 | ~ spl24_205), inference(avatar_split_clause, [], [f661, f1719, f1570, f1082, f1746])).
fof(f661, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f221])).
fof(f221, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f101])).
fof(f101, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e4 = op(e1, e1)) & (e3 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax101)).
fof(f1771, plain, (~ spl24_210 | ~ spl24_125 | ~ spl24_204 | ~ spl24_198), inference(avatar_split_clause, [], [f660, f1668, f1714, f1208, f1756])).
fof(f660, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e3 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f220])).
fof(f220, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e3 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f100])).
fof(f100, plain, ~ ((e2 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e4) & (e3 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax100)).
fof(f1770, plain, (~ spl24_208 | ~ spl24_65 | ~ spl24_179 | ~ spl24_197), inference(avatar_split_clause, [], [f659, f1663, f1564, f956, f1740])).
fof(f659, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f219])).
fof(f219, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f99])).
fof(f99, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), e2)) & (e4 = op(e2, e2)) & (e3 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax99)).
fof(f1769, plain, (~ spl24_209 | ~ spl24_95 | ~ spl24_203 | ~ spl24_182), inference(avatar_split_clause, [], [f658, f1580, f1708, f1082, f1746])).
fof(f658, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f218])).
fof(f218, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), e1)) & (e4 = op(e1, e1)) & (e3 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax98)).
fof(f1768, plain, (~ spl24_208 | ~ spl24_65 | ~ spl24_195 | ~ spl24_181), inference(avatar_split_clause, [], [f657, f1575, f1652, f956, f1740])).
fof(f657, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f217])).
fof(f217, plain, (~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f97])).
fof(f97, plain, ~ ((e1 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), e2)) & (e4 = op(e2, e2)) & (e3 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax97)).
fof(f1766, plain, (~ spl24_209 | ~ spl24_93 | ~ spl24_177 | ~ spl24_205), inference(avatar_split_clause, [], [f655, f1719, f1554, f1074, f1746])).
fof(f655, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f215])).
fof(f215, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f95])).
fof(f95, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1)) & (e3 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax95)).
fof(f1765, plain, (~ spl24_210 | ~ spl24_123 | ~ spl24_204 | ~ spl24_191), inference(avatar_split_clause, [], [f654, f1632, f1714, f1200, f1756])).
fof(f654, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f214])).
fof(f214, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f94])).
fof(f94, plain, ~ ((e4 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e2) & (e3 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax94)).
fof(f1763, plain, (~ spl24_209 | ~ spl24_93 | ~ spl24_203 | ~ spl24_175), inference(avatar_split_clause, [], [f652, f1544, f1708, f1074, f1746])).
fof(f652, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f212])).
fof(f212, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f92])).
fof(f92, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1)) & (e3 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax92)).
fof(f1760, plain, (~ spl24_208 | ~ spl24_62 | ~ spl24_168 | ~ spl24_197), inference(avatar_split_clause, [], [f649, f1663, f1512, f944, f1740])).
fof(f649, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f209])).
fof(f209, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f89])).
fof(f89, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e4 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e3 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax89)).
fof(f1753, plain, (~ spl24_208 | ~ spl24_62 | ~ spl24_195 | ~ spl24_165), inference(avatar_split_clause, [], [f646, f1498, f1652, f944, f1740])).
fof(f646, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f206])).
fof(f206, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f86])).
fof(f86, plain, ~ ((e4 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e3 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax86)).
fof(f1751, plain, (~ spl24_209 | ~ spl24_91 | ~ spl24_177 | ~ spl24_182), inference(avatar_split_clause, [], [f644, f1580, f1554, f1066, f1746])).
fof(f644, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f204])).
fof(f204, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f84])).
fof(f84, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e3 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax84)).
fof(f1749, plain, (~ spl24_209 | ~ spl24_91 | ~ spl24_180 | ~ spl24_175), inference(avatar_split_clause, [], [f642, f1544, f1570, f1066, f1746])).
fof(f642, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f202])).
fof(f202, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f82])).
fof(f82, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e3 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax82)).
fof(f1744, plain, (~ spl24_207 | ~ spl24_1 | ~ spl24_156 | ~ spl24_172), inference(avatar_split_clause, [], [f641, f1531, f1459, f688, f1735])).
fof(f641, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e3 = op(op(e4, e4), e4))), inference(cnf_transformation, [], [f201])).
fof(f201, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e3 = op(op(e4, e4), e4))), inference(ennf_transformation, [], [f81])).
fof(f81, plain, ~ ((e1 = op(e4, op(e4, e4))) & (e2 = op(op(e4, op(e4, e4)), e4)) & (e0 = op(e4, e4)) & (e3 = op(op(e4, e4), e4))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax81)).
fof(f1743, plain, (~ spl24_208 | ~ spl24_61 | ~ spl24_179 | ~ spl24_165), inference(avatar_split_clause, [], [f640, f1498, f1564, f940, f1740])).
fof(f640, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f200])).
fof(f200, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f80])).
fof(f80, plain, ~ ((e4 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, e2)) & (e3 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax80)).
fof(f1738, plain, (~ spl24_207 | ~ spl24_1 | ~ spl24_170 | ~ spl24_162), inference(avatar_split_clause, [], [f639, f1485, f1521, f688, f1735])).
fof(f639, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e3 = op(op(e4, e4), e4))), inference(cnf_transformation, [], [f199])).
fof(f199, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e3 = op(op(e4, e4), e4))), inference(ennf_transformation, [], [f79])).
fof(f79, plain, ~ ((e2 = op(e4, op(e4, e4))) & (e1 = op(op(e4, op(e4, e4)), e4)) & (e0 = op(e4, e4)) & (e3 = op(op(e4, e4), e4))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax79)).
fof(f1733, plain, (~ spl24_202 | ~ spl24_125 | ~ spl24_190 | ~ spl24_206), inference(avatar_split_clause, [], [f638, f1724, f1628, f1208, f1700])).
fof(f638, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e2 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f198])).
fof(f198, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e4) | ~ (e2 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f78])).
fof(f78, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e3 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e4) & (e2 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax78)).
fof(f1730, plain, (~ spl24_200 | ~ spl24_35 | ~ spl24_171 | ~ spl24_192), inference(avatar_split_clause, [], [f635, f1637, f1526, f830, f1684])).
fof(f635, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f195])).
fof(f195, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f75])).
fof(f75, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), e3)) & (e4 = op(e3, e3)) & (e2 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax75)).
fof(f1728, plain, (~ spl24_200 | ~ spl24_35 | ~ spl24_187 | ~ spl24_176), inference(avatar_split_clause, [], [f633, f1549, f1614, f830, f1684])).
fof(f633, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f193])).
fof(f193, plain, (~ (e1 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f73])).
fof(f73, plain, ~ ((e1 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), e3)) & (e4 = op(e3, e3)) & (e2 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax73)).
fof(f1727, plain, (~ spl24_202 | ~ spl24_124 | ~ spl24_193 | ~ spl24_206), inference(avatar_split_clause, [], [f632, f1724, f1642, f1204, f1700])).
fof(f632, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f192])).
fof(f192, plain, (~ (e1 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f72])).
fof(f72, plain, ~ ((e1 = op(e0, op(e0, e0))) & (e4 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e2 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax72)).
fof(f1722, plain, (~ spl24_201 | ~ spl24_94 | ~ spl24_177 | ~ spl24_205), inference(avatar_split_clause, [], [f631, f1719, f1554, f1078, f1690])).
fof(f631, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f191])).
fof(f191, plain, (~ (e0 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f71])).
fof(f71, plain, ~ ((e0 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1)) & (e2 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax71)).
fof(f1717, plain, (~ spl24_202 | ~ spl24_124 | ~ spl24_204 | ~ spl24_191), inference(avatar_split_clause, [], [f630, f1632, f1714, f1204, f1700])).
fof(f630, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f190])).
fof(f190, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f70])).
fof(f70, plain, ~ ((e4 = op(e0, op(e0, e0))) & (e1 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e2 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax70)).
fof(f1711, plain, (~ spl24_201 | ~ spl24_94 | ~ spl24_203 | ~ spl24_175), inference(avatar_split_clause, [], [f628, f1544, f1708, f1078, f1690])).
fof(f628, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f188])).
fof(f188, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f68])).
fof(f68, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e0 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1)) & (e2 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax68)).
fof(f1704, plain, (~ spl24_200 | ~ spl24_32 | ~ spl24_166 | ~ spl24_192), inference(avatar_split_clause, [], [f625, f1637, f1503, f818, f1684])).
fof(f625, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f185])).
fof(f185, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f65])).
fof(f65, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), e3)) & (e1 = op(e3, e3)) & (e2 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax65)).
fof(f1695, plain, (~ spl24_201 | ~ spl24_91 | ~ spl24_177 | ~ spl24_178), inference(avatar_split_clause, [], [f620, f1558, f1554, f1066, f1690])).
fof(f620, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f180])).
fof(f180, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e4 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f60])).
fof(f60, plain, ~ ((e3 = op(e1, op(e1, e1))) & (e4 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e2 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax60)).
fof(f1693, plain, (~ spl24_201 | ~ spl24_91 | ~ spl24_174 | ~ spl24_175), inference(avatar_split_clause, [], [f618, f1544, f1540, f1066, f1690])).
fof(f618, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f178])).
fof(f178, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f58])).
fof(f58, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e0 = op(e1, e1)) & (e2 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax58)).
fof(f1688, plain, (~ spl24_199 | ~ spl24_1 | ~ spl24_161 | ~ spl24_172), inference(avatar_split_clause, [], [f617, f1531, f1481, f688, f1679])).
fof(f617, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e2 = op(op(e4, e4), e4))), inference(cnf_transformation, [], [f177])).
fof(f177, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e2 = op(op(e4, e4), e4))), inference(ennf_transformation, [], [f57])).
fof(f57, plain, ~ ((e1 = op(e4, op(e4, e4))) & (e3 = op(op(e4, op(e4, e4)), e4)) & (e0 = op(e4, e4)) & (e2 = op(op(e4, e4), e4))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax57)).
fof(f1682, plain, (~ spl24_199 | ~ spl24_1 | ~ spl24_170 | ~ spl24_157), inference(avatar_split_clause, [], [f615, f1463, f1521, f688, f1679])).
fof(f615, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e2 = op(op(e4, e4), e4))), inference(cnf_transformation, [], [f175])).
fof(f175, plain, (~ (e3 = op(e4, op(e4, e4))) | ~ (e1 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e2 = op(op(e4, e4), e4))), inference(ennf_transformation, [], [f55])).
fof(f55, plain, ~ ((e3 = op(e4, op(e4, e4))) & (e1 = op(op(e4, op(e4, e4)), e4)) & (e0 = op(e4, e4)) & (e2 = op(op(e4, e4), e4))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax55)).
fof(f1676, plain, (~ spl24_185 | ~ spl24_65 | ~ spl24_164 | ~ spl24_197), inference(avatar_split_clause, [], [f613, f1663, f1494, f956, f1602])).
fof(f613, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f173])).
fof(f173, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f53])).
fof(f53, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e3 = op(op(e2, op(e2, e2)), e2)) & (e4 = op(e2, e2)) & (e1 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax53)).
fof(f1674, plain, (~ spl24_184 | ~ spl24_35 | ~ spl24_159 | ~ spl24_192), inference(avatar_split_clause, [], [f611, f1637, f1472, f830, f1596])).
fof(f611, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f171])).
fof(f171, plain, (~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f51])).
fof(f51, plain, ~ ((e0 = op(e3, op(e3, e3))) & (e2 = op(op(e3, op(e3, e3)), e3)) & (e4 = op(e3, e3)) & (e1 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax51)).
fof(f1673, plain, (~ spl24_185 | ~ spl24_65 | ~ spl24_195 | ~ spl24_169), inference(avatar_split_clause, [], [f610, f1516, f1652, f956, f1602])).
fof(f610, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f170])).
fof(f170, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f50])).
fof(f50, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e0 = op(op(e2, op(e2, e2)), e2)) & (e4 = op(e2, e2)) & (e1 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax50)).
fof(f1672, plain, (~ spl24_184 | ~ spl24_35 | ~ spl24_187 | ~ spl24_167), inference(avatar_split_clause, [], [f609, f1507, f1614, f830, f1596])).
fof(f609, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f169])).
fof(f169, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e0 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f49])).
fof(f49, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e0 = op(op(e3, op(e3, e3)), e3)) & (e4 = op(e3, e3)) & (e1 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax49)).
fof(f1671, plain, (~ spl24_189 | ~ spl24_124 | ~ spl24_193 | ~ spl24_198), inference(avatar_split_clause, [], [f608, f1668, f1642, f1204, f1624])).
fof(f608, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f168])).
fof(f168, plain, (~ (e2 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f48])).
fof(f48, plain, ~ ((e2 = op(e0, op(e0, e0))) & (e4 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e1 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax48)).
fof(f1666, plain, (~ spl24_185 | ~ spl24_64 | ~ spl24_168 | ~ spl24_197), inference(avatar_split_clause, [], [f607, f1663, f1512, f952, f1602])).
fof(f607, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (~ (e0 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f47])).
fof(f47, plain, ~ ((e0 = op(e2, op(e2, e2))) & (e4 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, e2)) & (e1 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax47)).
fof(f1661, plain, (~ spl24_189 | ~ spl24_124 | ~ spl24_196 | ~ spl24_191), inference(avatar_split_clause, [], [f606, f1632, f1658, f1204, f1624])).
fof(f606, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f166])).
fof(f166, plain, (~ (e4 = op(e0, op(e0, e0))) | ~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f46])).
fof(f46, plain, ~ ((e4 = op(e0, op(e0, e0))) & (e2 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e3) & (e1 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax46)).
fof(f1656, plain, (~ spl24_183 | ~ spl24_4 | ~ spl24_156 | ~ spl24_188), inference(avatar_split_clause, [], [f605, f1619, f1459, f700, f1591])).
fof(f605, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), e4)) | ~ (e3 = op(e4, e4)) | ~ (e1 = op(op(e4, e4), e4))), inference(cnf_transformation, [], [f165])).
fof(f165, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), e4)) | ~ (e3 = op(e4, e4)) | ~ (e1 = op(op(e4, e4), e4))), inference(ennf_transformation, [], [f45])).
fof(f45, plain, ~ ((e0 = op(e4, op(e4, e4))) & (e2 = op(op(e4, op(e4, e4)), e4)) & (e3 = op(e4, e4)) & (e1 = op(op(e4, e4), e4))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax45)).
fof(f1649, plain, (~ spl24_189 | ~ spl24_123 | ~ spl24_193 | ~ spl24_194), inference(avatar_split_clause, [], [f602, f1646, f1642, f1200, f1624])).
fof(f602, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0))), inference(cnf_transformation, [], [f162])).
fof(f162, plain, (~ (e3 = op(e0, op(e0, e0))) | ~ (e4 = op(op(e0, op(e0, e0)), e0)) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0))), inference(ennf_transformation, [], [f42])).
fof(f42, plain, ~ ((e3 = op(e0, op(e0, e0))) & (e4 = op(op(e0, op(e0, e0)), e0)) & (op(e0, e0) = e2) & (e1 = op(op(e0, e0), e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax42)).
fof(f1622, plain, (~ spl24_183 | ~ spl24_3 | ~ spl24_161 | ~ spl24_188), inference(avatar_split_clause, [], [f599, f1619, f1481, f696, f1591])).
fof(f599, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), e4)) | ~ (e2 = op(e4, e4)) | ~ (e1 = op(op(e4, e4), e4))), inference(cnf_transformation, [], [f159])).
fof(f159, plain, (~ (e0 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), e4)) | ~ (e2 = op(e4, e4)) | ~ (e1 = op(op(e4, e4), e4))), inference(ennf_transformation, [], [f39])).
fof(f39, plain, ~ ((e0 = op(e4, op(e4, e4))) & (e3 = op(op(e4, op(e4, e4)), e4)) & (e2 = op(e4, e4)) & (e1 = op(op(e4, e4), e4))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax39)).
fof(f1607, plain, (~ spl24_185 | ~ spl24_61 | ~ spl24_168 | ~ spl24_169), inference(avatar_split_clause, [], [f596, f1516, f1512, f940, f1602])).
fof(f596, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f156])).
fof(f156, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e4 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, e2)) & (e1 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax36)).
fof(f1606, plain, (~ spl24_184 | ~ spl24_31 | ~ spl24_166 | ~ spl24_167), inference(avatar_split_clause, [], [f595, f1507, f1503, f814, f1596])).
fof(f595, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f155])).
fof(f155, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e4 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e4 = op(op(e3, op(e3, e3)), e3)) & (e0 = op(e3, e3)) & (e1 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax35)).
fof(f1605, plain, (~ spl24_185 | ~ spl24_61 | ~ spl24_164 | ~ spl24_165), inference(avatar_split_clause, [], [f594, f1498, f1494, f940, f1602])).
fof(f594, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f154])).
fof(f154, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f34])).
fof(f34, plain, ~ ((e4 = op(e2, op(e2, e2))) & (e3 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, e2)) & (e1 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax34)).
fof(f1600, plain, (~ spl24_183 | ~ spl24_1 | ~ spl24_161 | ~ spl24_162), inference(avatar_split_clause, [], [f593, f1485, f1481, f688, f1591])).
fof(f593, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e1 = op(op(e4, e4), e4))), inference(cnf_transformation, [], [f153])).
fof(f153, plain, (~ (e2 = op(e4, op(e4, e4))) | ~ (e3 = op(op(e4, op(e4, e4)), e4)) | ~ (e0 = op(e4, e4)) | ~ (e1 = op(op(e4, e4), e4))), inference(ennf_transformation, [], [f33])).
fof(f33, plain, ~ ((e2 = op(e4, op(e4, e4))) & (e3 = op(op(e4, op(e4, e4)), e4)) & (e0 = op(e4, e4)) & (e1 = op(op(e4, e4), e4))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax33)).
fof(f1589, plain, (~ spl24_173 | ~ spl24_95 | ~ spl24_174 | ~ spl24_182), inference(avatar_split_clause, [], [f590, f1580, f1540, f1082, f1536])).
fof(f590, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f150])).
fof(f150, plain, (~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f30])).
fof(f30, plain, ~ ((e2 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e4 = op(e1, e1)) & (e0 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax30)).
fof(f1587, plain, (~ spl24_173 | ~ spl24_95 | ~ spl24_180 | ~ spl24_178), inference(avatar_split_clause, [], [f588, f1558, f1570, f1082, f1536])).
fof(f588, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f148])).
fof(f148, plain, (~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e4 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e3 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e4 = op(e1, e1)) & (e0 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax28)).
fof(f1585, plain, (~ spl24_163 | ~ spl24_65 | ~ spl24_179 | ~ spl24_169), inference(avatar_split_clause, [], [f586, f1516, f1564, f956, f1490])).
fof(f586, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f146])).
fof(f146, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e4 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e1 = op(op(e2, op(e2, e2)), e2)) & (e4 = op(e2, e2)) & (e0 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax26)).
fof(f1584, plain, (~ spl24_158 | ~ spl24_35 | ~ spl24_171 | ~ spl24_167), inference(avatar_split_clause, [], [f585, f1507, f1526, f830, f1468])).
fof(f585, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3))), inference(cnf_transformation, [], [f145])).
fof(f145, plain, (~ (e2 = op(e3, op(e3, e3))) | ~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e4 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e2 = op(e3, op(e3, e3))) & (e1 = op(op(e3, op(e3, e3)), e3)) & (e4 = op(e3, e3)) & (e0 = op(op(e3, e3), e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax25)).
fof(f1573, plain, (~ spl24_173 | ~ spl24_94 | ~ spl24_180 | ~ spl24_175), inference(avatar_split_clause, [], [f582, f1544, f1570, f1078, f1536])).
fof(f582, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f142])).
fof(f142, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e2 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e2 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, e1)) & (e0 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax22)).
fof(f1568, plain, (~ spl24_155 | ~ spl24_4 | ~ spl24_156 | ~ spl24_172), inference(avatar_split_clause, [], [f581, f1531, f1459, f700, f1455])).
fof(f581, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), e4)) | ~ (e3 = op(e4, e4)) | ~ (e0 = op(op(e4, e4), e4))), inference(cnf_transformation, [], [f141])).
fof(f141, plain, (~ (e1 = op(e4, op(e4, e4))) | ~ (e2 = op(op(e4, op(e4, e4)), e4)) | ~ (e3 = op(e4, e4)) | ~ (e0 = op(op(e4, e4), e4))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e1 = op(e4, op(e4, e4))) & (e2 = op(op(e4, op(e4, e4)), e4)) & (e3 = op(e4, e4)) & (e0 = op(op(e4, e4), e4))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax21)).
fof(f1547, plain, (~ spl24_173 | ~ spl24_93 | ~ spl24_174 | ~ spl24_175), inference(avatar_split_clause, [], [f576, f1544, f1540, f1074, f1536])).
fof(f576, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(cnf_transformation, [], [f136])).
fof(f136, plain, (~ (e4 = op(e1, op(e1, e1))) | ~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e4 = op(e1, op(e1, e1))) & (e3 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, e1)) & (e0 = op(op(e1, e1), e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax16)).
fof(f1519, plain, (~ spl24_163 | ~ spl24_62 | ~ spl24_168 | ~ spl24_169), inference(avatar_split_clause, [], [f572, f1516, f1512, f944, f1490])).
fof(f572, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f132])).
fof(f132, plain, (~ (e3 = op(e2, op(e2, e2))) | ~ (e4 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(e2, op(e2, e2))) & (e4 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e0 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax12)).
fof(f1501, plain, (~ spl24_163 | ~ spl24_62 | ~ spl24_164 | ~ spl24_165), inference(avatar_split_clause, [], [f570, f1498, f1494, f944, f1490])).
fof(f570, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2))), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (~ (e4 = op(e2, op(e2, e2))) | ~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e4 = op(e2, op(e2, e2))) & (e3 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, e2)) & (e0 = op(op(e2, e2), e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax10)).
fof(f1453, plain, (spl24_154 | spl24_153 | spl24_152 | spl24_151 | spl24_150 | spl24_149 | spl24_148 | spl24_147 | spl24_146 | spl24_145 | spl24_144 | spl24_143 | spl24_142 | spl24_141 | spl24_140 | spl24_139 | spl24_138 | spl24_137 | spl24_136 | spl24_135 | spl24_134 | spl24_133 | spl24_132 | spl24_131 | spl24_5), inference(avatar_split_clause, [], [f564, f704, f1284, f1291, f1298, f1305, f1312, f1319, f1326, f1333, f1340, f1347, f1354, f1361, f1368, f1375, f1382, f1389, f1396, f1403, f1410, f1417, f1424, f1431, f1438, f1445])).
fof(f1445, plain, (spl24_154 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl24_154])])).
fof(f1438, plain, (spl24_153 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl24_153])])).
fof(f1431, plain, (spl24_152 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl24_152])])).
fof(f1424, plain, (spl24_151 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl24_151])])).
fof(f1417, plain, (spl24_150 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl24_150])])).
fof(f1410, plain, (spl24_149 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl24_149])])).
fof(f1403, plain, (spl24_148 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl24_148])])).
fof(f1396, plain, (spl24_147 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl24_147])])).
fof(f1389, plain, (spl24_146 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl24_146])])).
fof(f1382, plain, (spl24_145 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl24_145])])).
fof(f1375, plain, (spl24_144 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl24_144])])).
fof(f1368, plain, (spl24_143 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl24_143])])).
fof(f1361, plain, (spl24_142 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl24_142])])).
fof(f1354, plain, (spl24_141 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl24_141])])).
fof(f1347, plain, (spl24_140 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl24_140])])).
fof(f1340, plain, (spl24_139 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl24_139])])).
fof(f1333, plain, (spl24_138 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl24_138])])).
fof(f1326, plain, (spl24_137 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl24_137])])).
fof(f1319, plain, (spl24_136 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl24_136])])).
fof(f1312, plain, (spl24_135 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl24_135])])).
fof(f1305, plain, (spl24_134 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl24_134])])).
fof(f1298, plain, (spl24_133 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl24_133])])).
fof(f1291, plain, (spl24_132 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl24_132])])).
fof(f1284, plain, (spl24_131 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl24_131])])).
fof(f564, plain, ((e4 = op(e4, e4)) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ((~ (e4 = op(e4, e4)) & (e4 = op(e4, e4)) & (e4 = op(e4, e4))) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_folding, [], [f6, e270, e269, e268, e267, e266, e265, e264, e263, e262, e261, e260, e259, e258, e257, e256, e255, e254, e253, e252, e251, e250, e249, e248, e247])).
fof(f247, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e247])).
fof(e247, plain, (sP0 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f248, plain, ((~ (e0 = op(e0, e1)) & (e0 = op(e1, e1)) & (op(e0, e0) = e1)) | ~ sP1), inference(usedef, [], [e248])).
fof(e248, plain, (sP1 <=> (~ (e0 = op(e0, e1)) & (e0 = op(e1, e1)) & (op(e0, e0) = e1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f249, plain, ((~ (e0 = op(e0, e2)) & (e0 = op(e2, e2)) & (op(e0, e0) = e2)) | ~ sP2), inference(usedef, [], [e249])).
fof(e249, plain, (sP2 <=> (~ (e0 = op(e0, e2)) & (e0 = op(e2, e2)) & (op(e0, e0) = e2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f250, plain, ((~ (e0 = op(e0, e3)) & (e0 = op(e3, e3)) & (op(e0, e0) = e3)) | ~ sP3), inference(usedef, [], [e250])).
fof(e250, plain, (sP3 <=> (~ (e0 = op(e0, e3)) & (e0 = op(e3, e3)) & (op(e0, e0) = e3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f251, plain, ((~ (e0 = op(e0, e4)) & (e0 = op(e4, e4)) & (op(e0, e0) = e4)) | ~ sP4), inference(usedef, [], [e251])).
fof(e251, plain, (sP4 <=> (~ (e0 = op(e0, e4)) & (e0 = op(e4, e4)) & (op(e0, e0) = e4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f252, plain, ((~ (e1 = op(e1, e0)) & (op(e0, e0) = e1) & (e0 = op(e1, e1))) | ~ sP5), inference(usedef, [], [e252])).
fof(e252, plain, (sP5 <=> (~ (e1 = op(e1, e0)) & (op(e0, e0) = e1) & (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f253, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP6), inference(usedef, [], [e253])).
fof(e253, plain, (sP6 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f254, plain, ((~ (e1 = op(e1, e2)) & (e1 = op(e2, e2)) & (e2 = op(e1, e1))) | ~ sP7), inference(usedef, [], [e254])).
fof(e254, plain, (sP7 <=> (~ (e1 = op(e1, e2)) & (e1 = op(e2, e2)) & (e2 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f255, plain, ((~ (e1 = op(e1, e3)) & (e1 = op(e3, e3)) & (e3 = op(e1, e1))) | ~ sP8), inference(usedef, [], [e255])).
fof(e255, plain, (sP8 <=> (~ (e1 = op(e1, e3)) & (e1 = op(e3, e3)) & (e3 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f256, plain, ((~ (e1 = op(e1, e4)) & (e1 = op(e4, e4)) & (e4 = op(e1, e1))) | ~ sP9), inference(usedef, [], [e256])).
fof(e256, plain, (sP9 <=> (~ (e1 = op(e1, e4)) & (e1 = op(e4, e4)) & (e4 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f257, plain, ((~ (e2 = op(e2, e0)) & (op(e0, e0) = e2) & (e0 = op(e2, e2))) | ~ sP10), inference(usedef, [], [e257])).
fof(e257, plain, (sP10 <=> (~ (e2 = op(e2, e0)) & (op(e0, e0) = e2) & (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f258, plain, ((~ (e2 = op(e2, e1)) & (e2 = op(e1, e1)) & (e1 = op(e2, e2))) | ~ sP11), inference(usedef, [], [e258])).
fof(e258, plain, (sP11 <=> (~ (e2 = op(e2, e1)) & (e2 = op(e1, e1)) & (e1 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f259, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP12), inference(usedef, [], [e259])).
fof(e259, plain, (sP12 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f260, plain, ((~ (e2 = op(e2, e3)) & (e2 = op(e3, e3)) & (e3 = op(e2, e2))) | ~ sP13), inference(usedef, [], [e260])).
fof(e260, plain, (sP13 <=> (~ (e2 = op(e2, e3)) & (e2 = op(e3, e3)) & (e3 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f261, plain, ((~ (e2 = op(e2, e4)) & (e2 = op(e4, e4)) & (e4 = op(e2, e2))) | ~ sP14), inference(usedef, [], [e261])).
fof(e261, plain, (sP14 <=> (~ (e2 = op(e2, e4)) & (e2 = op(e4, e4)) & (e4 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f262, plain, ((~ (e3 = op(e3, e0)) & (op(e0, e0) = e3) & (e0 = op(e3, e3))) | ~ sP15), inference(usedef, [], [e262])).
fof(e262, plain, (sP15 <=> (~ (e3 = op(e3, e0)) & (op(e0, e0) = e3) & (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f263, plain, ((~ (e3 = op(e3, e1)) & (e3 = op(e1, e1)) & (e1 = op(e3, e3))) | ~ sP16), inference(usedef, [], [e263])).
fof(e263, plain, (sP16 <=> (~ (e3 = op(e3, e1)) & (e3 = op(e1, e1)) & (e1 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f264, plain, ((~ (e3 = op(e3, e2)) & (e3 = op(e2, e2)) & (e2 = op(e3, e3))) | ~ sP17), inference(usedef, [], [e264])).
fof(e264, plain, (sP17 <=> (~ (e3 = op(e3, e2)) & (e3 = op(e2, e2)) & (e2 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f265, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP18), inference(usedef, [], [e265])).
fof(e265, plain, (sP18 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f266, plain, ((~ (e3 = op(e3, e4)) & (e3 = op(e4, e4)) & (e4 = op(e3, e3))) | ~ sP19), inference(usedef, [], [e266])).
fof(e266, plain, (sP19 <=> (~ (e3 = op(e3, e4)) & (e3 = op(e4, e4)) & (e4 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f267, plain, ((~ (e4 = op(e4, e0)) & (op(e0, e0) = e4) & (e0 = op(e4, e4))) | ~ sP20), inference(usedef, [], [e267])).
fof(e267, plain, (sP20 <=> (~ (e4 = op(e4, e0)) & (op(e0, e0) = e4) & (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f268, plain, ((~ (e4 = op(e4, e1)) & (e4 = op(e1, e1)) & (e1 = op(e4, e4))) | ~ sP21), inference(usedef, [], [e268])).
fof(e268, plain, (sP21 <=> (~ (e4 = op(e4, e1)) & (e4 = op(e1, e1)) & (e1 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f269, plain, ((~ (e4 = op(e4, e2)) & (e4 = op(e2, e2)) & (e2 = op(e4, e4))) | ~ sP22), inference(usedef, [], [e269])).
fof(e269, plain, (sP22 <=> (~ (e4 = op(e4, e2)) & (e4 = op(e2, e2)) & (e2 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f270, plain, ((~ (e4 = op(e4, e3)) & (e4 = op(e3, e3)) & (e3 = op(e4, e4))) | ~ sP23), inference(usedef, [], [e270])).
fof(e270, plain, (sP23 <=> (~ (e4 = op(e4, e3)) & (e4 = op(e3, e3)) & (e3 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f6, plain, ((~ (e4 = op(e4, e4)) & (e4 = op(e4, e4)) & (e4 = op(e4, e4))) | (~ (e4 = op(e4, e3)) & (e4 = op(e3, e3)) & (e3 = op(e4, e4))) | (~ (e4 = op(e4, e2)) & (e4 = op(e2, e2)) & (e2 = op(e4, e4))) | (~ (e4 = op(e4, e1)) & (e4 = op(e1, e1)) & (e1 = op(e4, e4))) | (~ (e4 = op(e4, e0)) & (op(e0, e0) = e4) & (e0 = op(e4, e4))) | (~ (e3 = op(e3, e4)) & (e3 = op(e4, e4)) & (e4 = op(e3, e3))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e3 = op(e2, e2)) & (e2 = op(e3, e3))) | (~ (e3 = op(e3, e1)) & (e3 = op(e1, e1)) & (e1 = op(e3, e3))) | (~ (e3 = op(e3, e0)) & (op(e0, e0) = e3) & (e0 = op(e3, e3))) | (~ (e2 = op(e2, e4)) & (e2 = op(e4, e4)) & (e4 = op(e2, e2))) | (~ (e2 = op(e2, e3)) & (e2 = op(e3, e3)) & (e3 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e1)) & (e2 = op(e1, e1)) & (e1 = op(e2, e2))) | (~ (e2 = op(e2, e0)) & (op(e0, e0) = e2) & (e0 = op(e2, e2))) | (~ (e1 = op(e1, e4)) & (e1 = op(e4, e4)) & (e4 = op(e1, e1))) | (~ (e1 = op(e1, e3)) & (e1 = op(e3, e3)) & (e3 = op(e1, e1))) | (~ (e1 = op(e1, e2)) & (e1 = op(e2, e2)) & (e2 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e0)) & (op(e0, e0) = e1) & (e0 = op(e1, e1))) | (~ (e0 = op(e0, e4)) & (e0 = op(e4, e4)) & (op(e0, e0) = e4)) | (~ (e0 = op(e0, e3)) & (e0 = op(e3, e3)) & (op(e0, e0) = e3)) | (~ (e0 = op(e0, e2)) & (e0 = op(e2, e2)) & (op(e0, e0) = e2)) | (~ (e0 = op(e0, e1)) & (e0 = op(e1, e1)) & (op(e0, e0) = e1)) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax6)).
fof(f1451, plain, (spl24_154 | spl24_153 | spl24_152 | spl24_151 | spl24_150 | spl24_149 | spl24_148 | spl24_147 | spl24_146 | spl24_145 | spl24_144 | spl24_143 | spl24_142 | spl24_141 | spl24_140 | spl24_139 | spl24_138 | spl24_137 | spl24_136 | spl24_135 | spl24_134 | spl24_133 | spl24_132 | spl24_131 | ~ spl24_5), inference(avatar_split_clause, [], [f566, f704, f1284, f1291, f1298, f1305, f1312, f1319, f1326, f1333, f1340, f1347, f1354, f1361, f1368, f1375, f1382, f1389, f1396, f1403, f1410, f1417, f1424, f1431, f1438, f1445])).
fof(f566, plain, (~ (e4 = op(e4, e4)) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f271])).
fof(f1450, plain, (~ spl24_154 | spl24_121), inference(avatar_split_clause, [], [f561, f1192, f1445])).
fof(f561, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f295])).
fof(f295, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f247])).
fof(f1448, plain, (~ spl24_154 | ~ spl24_121), inference(avatar_split_clause, [], [f563, f1192, f1445])).
fof(f563, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f295])).
fof(f1443, plain, (~ spl24_153 | spl24_122), inference(avatar_split_clause, [], [f558, f1196, f1438])).
fof(f558, plain, ((op(e0, e0) = e1) | ~ sP1), inference(cnf_transformation, [], [f294])).
fof(f294, plain, ((~ (e0 = op(e0, e1)) & (e0 = op(e1, e1)) & (op(e0, e0) = e1)) | ~ sP1), inference(nnf_transformation, [], [f248])).
fof(f1442, plain, (~ spl24_153 | spl24_91), inference(avatar_split_clause, [], [f559, f1066, f1438])).
fof(f559, plain, ((e0 = op(e1, e1)) | ~ sP1), inference(cnf_transformation, [], [f294])).
fof(f1441, plain, (~ spl24_153 | ~ spl24_116), inference(avatar_split_clause, [], [f560, f1171, f1438])).
fof(f560, plain, (~ (e0 = op(e0, e1)) | ~ sP1), inference(cnf_transformation, [], [f294])).
fof(f1436, plain, (~ spl24_152 | spl24_123), inference(avatar_split_clause, [], [f555, f1200, f1431])).
fof(f555, plain, ((op(e0, e0) = e2) | ~ sP2), inference(cnf_transformation, [], [f293])).
fof(f293, plain, ((~ (e0 = op(e0, e2)) & (e0 = op(e2, e2)) & (op(e0, e0) = e2)) | ~ sP2), inference(nnf_transformation, [], [f249])).
fof(f1435, plain, (~ spl24_152 | spl24_61), inference(avatar_split_clause, [], [f556, f940, f1431])).
fof(f556, plain, ((e0 = op(e2, e2)) | ~ sP2), inference(cnf_transformation, [], [f293])).
fof(f1434, plain, (~ spl24_152 | ~ spl24_111), inference(avatar_split_clause, [], [f557, f1150, f1431])).
fof(f557, plain, (~ (e0 = op(e0, e2)) | ~ sP2), inference(cnf_transformation, [], [f293])).
fof(f1429, plain, (~ spl24_151 | spl24_124), inference(avatar_split_clause, [], [f552, f1204, f1424])).
fof(f552, plain, ((op(e0, e0) = e3) | ~ sP3), inference(cnf_transformation, [], [f292])).
fof(f292, plain, ((~ (e0 = op(e0, e3)) & (e0 = op(e3, e3)) & (op(e0, e0) = e3)) | ~ sP3), inference(nnf_transformation, [], [f250])).
fof(f1428, plain, (~ spl24_151 | spl24_31), inference(avatar_split_clause, [], [f553, f814, f1424])).
fof(f553, plain, ((e0 = op(e3, e3)) | ~ sP3), inference(cnf_transformation, [], [f292])).
fof(f1427, plain, (~ spl24_151 | ~ spl24_106), inference(avatar_split_clause, [], [f554, f1129, f1424])).
fof(f554, plain, (~ (e0 = op(e0, e3)) | ~ sP3), inference(cnf_transformation, [], [f292])).
fof(f1422, plain, (~ spl24_150 | spl24_125), inference(avatar_split_clause, [], [f549, f1208, f1417])).
fof(f549, plain, ((op(e0, e0) = e4) | ~ sP4), inference(cnf_transformation, [], [f291])).
fof(f291, plain, ((~ (e0 = op(e0, e4)) & (e0 = op(e4, e4)) & (op(e0, e0) = e4)) | ~ sP4), inference(nnf_transformation, [], [f251])).
fof(f1421, plain, (~ spl24_150 | spl24_1), inference(avatar_split_clause, [], [f550, f688, f1417])).
fof(f550, plain, ((e0 = op(e4, e4)) | ~ sP4), inference(cnf_transformation, [], [f291])).
fof(f1420, plain, (~ spl24_150 | ~ spl24_101), inference(avatar_split_clause, [], [f551, f1108, f1417])).
fof(f551, plain, (~ (e0 = op(e0, e4)) | ~ sP4), inference(cnf_transformation, [], [f291])).
fof(f1415, plain, (~ spl24_149 | spl24_91), inference(avatar_split_clause, [], [f546, f1066, f1410])).
fof(f546, plain, ((e0 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f290])).
fof(f290, plain, ((~ (e1 = op(e1, e0)) & (op(e0, e0) = e1) & (e0 = op(e1, e1))) | ~ sP5), inference(nnf_transformation, [], [f252])).
fof(f1414, plain, (~ spl24_149 | spl24_122), inference(avatar_split_clause, [], [f547, f1196, f1410])).
fof(f547, plain, ((op(e0, e0) = e1) | ~ sP5), inference(cnf_transformation, [], [f290])).
fof(f1413, plain, (~ spl24_149 | ~ spl24_97), inference(avatar_split_clause, [], [f548, f1091, f1410])).
fof(f548, plain, (~ (e1 = op(e1, e0)) | ~ sP5), inference(cnf_transformation, [], [f290])).
fof(f1408, plain, (~ spl24_148 | spl24_92), inference(avatar_split_clause, [], [f543, f1070, f1403])).
fof(f543, plain, ((e1 = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f289])).
fof(f289, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP6), inference(nnf_transformation, [], [f253])).
fof(f1406, plain, (~ spl24_148 | ~ spl24_92), inference(avatar_split_clause, [], [f545, f1070, f1403])).
fof(f545, plain, (~ (e1 = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f289])).
fof(f1401, plain, (~ spl24_147 | spl24_93), inference(avatar_split_clause, [], [f540, f1074, f1396])).
fof(f540, plain, ((e2 = op(e1, e1)) | ~ sP7), inference(cnf_transformation, [], [f288])).
fof(f288, plain, ((~ (e1 = op(e1, e2)) & (e1 = op(e2, e2)) & (e2 = op(e1, e1))) | ~ sP7), inference(nnf_transformation, [], [f254])).
fof(f1400, plain, (~ spl24_147 | spl24_62), inference(avatar_split_clause, [], [f541, f944, f1396])).
fof(f541, plain, ((e1 = op(e2, e2)) | ~ sP7), inference(cnf_transformation, [], [f288])).
fof(f1399, plain, (~ spl24_147 | ~ spl24_87), inference(avatar_split_clause, [], [f542, f1049, f1396])).
fof(f542, plain, (~ (e1 = op(e1, e2)) | ~ sP7), inference(cnf_transformation, [], [f288])).
fof(f1394, plain, (~ spl24_146 | spl24_94), inference(avatar_split_clause, [], [f537, f1078, f1389])).
fof(f537, plain, ((e3 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f287])).
fof(f287, plain, ((~ (e1 = op(e1, e3)) & (e1 = op(e3, e3)) & (e3 = op(e1, e1))) | ~ sP8), inference(nnf_transformation, [], [f255])).
fof(f1393, plain, (~ spl24_146 | spl24_32), inference(avatar_split_clause, [], [f538, f818, f1389])).
fof(f538, plain, ((e1 = op(e3, e3)) | ~ sP8), inference(cnf_transformation, [], [f287])).
fof(f1392, plain, (~ spl24_146 | ~ spl24_82), inference(avatar_split_clause, [], [f539, f1028, f1389])).
fof(f539, plain, (~ (e1 = op(e1, e3)) | ~ sP8), inference(cnf_transformation, [], [f287])).
fof(f1387, plain, (~ spl24_145 | spl24_95), inference(avatar_split_clause, [], [f534, f1082, f1382])).
fof(f534, plain, ((e4 = op(e1, e1)) | ~ sP9), inference(cnf_transformation, [], [f286])).
fof(f286, plain, ((~ (e1 = op(e1, e4)) & (e1 = op(e4, e4)) & (e4 = op(e1, e1))) | ~ sP9), inference(nnf_transformation, [], [f256])).
fof(f1386, plain, (~ spl24_145 | spl24_2), inference(avatar_split_clause, [], [f535, f692, f1382])).
fof(f535, plain, ((e1 = op(e4, e4)) | ~ sP9), inference(cnf_transformation, [], [f286])).
fof(f1385, plain, (~ spl24_145 | ~ spl24_77), inference(avatar_split_clause, [], [f536, f1007, f1382])).
fof(f536, plain, (~ (e1 = op(e1, e4)) | ~ sP9), inference(cnf_transformation, [], [f286])).
fof(f1380, plain, (~ spl24_144 | spl24_61), inference(avatar_split_clause, [], [f531, f940, f1375])).
fof(f531, plain, ((e0 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f285])).
fof(f285, plain, ((~ (e2 = op(e2, e0)) & (op(e0, e0) = e2) & (e0 = op(e2, e2))) | ~ sP10), inference(nnf_transformation, [], [f257])).
fof(f1379, plain, (~ spl24_144 | spl24_123), inference(avatar_split_clause, [], [f532, f1200, f1375])).
fof(f532, plain, ((op(e0, e0) = e2) | ~ sP10), inference(cnf_transformation, [], [f285])).
fof(f1378, plain, (~ spl24_144 | ~ spl24_73), inference(avatar_split_clause, [], [f533, f990, f1375])).
fof(f533, plain, (~ (e2 = op(e2, e0)) | ~ sP10), inference(cnf_transformation, [], [f285])).
fof(f1373, plain, (~ spl24_143 | spl24_62), inference(avatar_split_clause, [], [f528, f944, f1368])).
fof(f528, plain, ((e1 = op(e2, e2)) | ~ sP11), inference(cnf_transformation, [], [f284])).
fof(f284, plain, ((~ (e2 = op(e2, e1)) & (e2 = op(e1, e1)) & (e1 = op(e2, e2))) | ~ sP11), inference(nnf_transformation, [], [f258])).
fof(f1372, plain, (~ spl24_143 | spl24_93), inference(avatar_split_clause, [], [f529, f1074, f1368])).
fof(f529, plain, ((e2 = op(e1, e1)) | ~ sP11), inference(cnf_transformation, [], [f284])).
fof(f1371, plain, (~ spl24_143 | ~ spl24_68), inference(avatar_split_clause, [], [f530, f969, f1368])).
fof(f530, plain, (~ (e2 = op(e2, e1)) | ~ sP11), inference(cnf_transformation, [], [f284])).
fof(f1366, plain, (~ spl24_142 | spl24_63), inference(avatar_split_clause, [], [f525, f948, f1361])).
fof(f525, plain, ((e2 = op(e2, e2)) | ~ sP12), inference(cnf_transformation, [], [f283])).
fof(f283, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP12), inference(nnf_transformation, [], [f259])).
fof(f1364, plain, (~ spl24_142 | ~ spl24_63), inference(avatar_split_clause, [], [f527, f948, f1361])).
fof(f527, plain, (~ (e2 = op(e2, e2)) | ~ sP12), inference(cnf_transformation, [], [f283])).
fof(f1359, plain, (~ spl24_141 | spl24_64), inference(avatar_split_clause, [], [f522, f952, f1354])).
fof(f522, plain, ((e3 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f282])).
fof(f282, plain, ((~ (e2 = op(e2, e3)) & (e2 = op(e3, e3)) & (e3 = op(e2, e2))) | ~ sP13), inference(nnf_transformation, [], [f260])).
fof(f1358, plain, (~ spl24_141 | spl24_33), inference(avatar_split_clause, [], [f523, f822, f1354])).
fof(f523, plain, ((e2 = op(e3, e3)) | ~ sP13), inference(cnf_transformation, [], [f282])).
fof(f1357, plain, (~ spl24_141 | ~ spl24_58), inference(avatar_split_clause, [], [f524, f927, f1354])).
fof(f524, plain, (~ (e2 = op(e2, e3)) | ~ sP13), inference(cnf_transformation, [], [f282])).
fof(f1352, plain, (~ spl24_140 | spl24_65), inference(avatar_split_clause, [], [f519, f956, f1347])).
fof(f519, plain, ((e4 = op(e2, e2)) | ~ sP14), inference(cnf_transformation, [], [f281])).
fof(f281, plain, ((~ (e2 = op(e2, e4)) & (e2 = op(e4, e4)) & (e4 = op(e2, e2))) | ~ sP14), inference(nnf_transformation, [], [f261])).
fof(f1351, plain, (~ spl24_140 | spl24_3), inference(avatar_split_clause, [], [f520, f696, f1347])).
fof(f520, plain, ((e2 = op(e4, e4)) | ~ sP14), inference(cnf_transformation, [], [f281])).
fof(f1350, plain, (~ spl24_140 | ~ spl24_53), inference(avatar_split_clause, [], [f521, f906, f1347])).
fof(f521, plain, (~ (e2 = op(e2, e4)) | ~ sP14), inference(cnf_transformation, [], [f281])).
fof(f1345, plain, (~ spl24_139 | spl24_31), inference(avatar_split_clause, [], [f516, f814, f1340])).
fof(f516, plain, ((e0 = op(e3, e3)) | ~ sP15), inference(cnf_transformation, [], [f280])).
fof(f280, plain, ((~ (e3 = op(e3, e0)) & (op(e0, e0) = e3) & (e0 = op(e3, e3))) | ~ sP15), inference(nnf_transformation, [], [f262])).
fof(f1344, plain, (~ spl24_139 | spl24_124), inference(avatar_split_clause, [], [f517, f1204, f1340])).
fof(f517, plain, ((op(e0, e0) = e3) | ~ sP15), inference(cnf_transformation, [], [f280])).
fof(f1343, plain, (~ spl24_139 | ~ spl24_49), inference(avatar_split_clause, [], [f518, f889, f1340])).
fof(f518, plain, (~ (e3 = op(e3, e0)) | ~ sP15), inference(cnf_transformation, [], [f280])).
fof(f1338, plain, (~ spl24_138 | spl24_32), inference(avatar_split_clause, [], [f513, f818, f1333])).
fof(f513, plain, ((e1 = op(e3, e3)) | ~ sP16), inference(cnf_transformation, [], [f279])).
fof(f279, plain, ((~ (e3 = op(e3, e1)) & (e3 = op(e1, e1)) & (e1 = op(e3, e3))) | ~ sP16), inference(nnf_transformation, [], [f263])).
fof(f1337, plain, (~ spl24_138 | spl24_94), inference(avatar_split_clause, [], [f514, f1078, f1333])).
fof(f514, plain, ((e3 = op(e1, e1)) | ~ sP16), inference(cnf_transformation, [], [f279])).
fof(f1336, plain, (~ spl24_138 | ~ spl24_44), inference(avatar_split_clause, [], [f515, f868, f1333])).
fof(f515, plain, (~ (e3 = op(e3, e1)) | ~ sP16), inference(cnf_transformation, [], [f279])).
fof(f1331, plain, (~ spl24_137 | spl24_33), inference(avatar_split_clause, [], [f510, f822, f1326])).
fof(f510, plain, ((e2 = op(e3, e3)) | ~ sP17), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ((~ (e3 = op(e3, e2)) & (e3 = op(e2, e2)) & (e2 = op(e3, e3))) | ~ sP17), inference(nnf_transformation, [], [f264])).
fof(f1330, plain, (~ spl24_137 | spl24_64), inference(avatar_split_clause, [], [f511, f952, f1326])).
fof(f511, plain, ((e3 = op(e2, e2)) | ~ sP17), inference(cnf_transformation, [], [f278])).
fof(f1329, plain, (~ spl24_137 | ~ spl24_39), inference(avatar_split_clause, [], [f512, f847, f1326])).
fof(f512, plain, (~ (e3 = op(e3, e2)) | ~ sP17), inference(cnf_transformation, [], [f278])).
fof(f1324, plain, (~ spl24_136 | spl24_34), inference(avatar_split_clause, [], [f507, f826, f1319])).
fof(f507, plain, ((e3 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f277])).
fof(f277, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP18), inference(nnf_transformation, [], [f265])).
fof(f1322, plain, (~ spl24_136 | ~ spl24_34), inference(avatar_split_clause, [], [f509, f826, f1319])).
fof(f509, plain, (~ (e3 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f277])).
fof(f1317, plain, (~ spl24_135 | spl24_35), inference(avatar_split_clause, [], [f504, f830, f1312])).
fof(f504, plain, ((e4 = op(e3, e3)) | ~ sP19), inference(cnf_transformation, [], [f276])).
fof(f276, plain, ((~ (e3 = op(e3, e4)) & (e3 = op(e4, e4)) & (e4 = op(e3, e3))) | ~ sP19), inference(nnf_transformation, [], [f266])).
fof(f1316, plain, (~ spl24_135 | spl24_4), inference(avatar_split_clause, [], [f505, f700, f1312])).
fof(f505, plain, ((e3 = op(e4, e4)) | ~ sP19), inference(cnf_transformation, [], [f276])).
fof(f1315, plain, (~ spl24_135 | ~ spl24_29), inference(avatar_split_clause, [], [f506, f805, f1312])).
fof(f506, plain, (~ (e3 = op(e3, e4)) | ~ sP19), inference(cnf_transformation, [], [f276])).
fof(f1310, plain, (~ spl24_134 | spl24_1), inference(avatar_split_clause, [], [f501, f688, f1305])).
fof(f501, plain, ((e0 = op(e4, e4)) | ~ sP20), inference(cnf_transformation, [], [f275])).
fof(f275, plain, ((~ (e4 = op(e4, e0)) & (op(e0, e0) = e4) & (e0 = op(e4, e4))) | ~ sP20), inference(nnf_transformation, [], [f267])).
fof(f1309, plain, (~ spl24_134 | spl24_125), inference(avatar_split_clause, [], [f502, f1208, f1305])).
fof(f502, plain, ((op(e0, e0) = e4) | ~ sP20), inference(cnf_transformation, [], [f275])).
fof(f1303, plain, (~ spl24_133 | spl24_2), inference(avatar_split_clause, [], [f498, f692, f1298])).
fof(f498, plain, ((e1 = op(e4, e4)) | ~ sP21), inference(cnf_transformation, [], [f274])).
fof(f274, plain, ((~ (e4 = op(e4, e1)) & (e4 = op(e1, e1)) & (e1 = op(e4, e4))) | ~ sP21), inference(nnf_transformation, [], [f268])).
fof(f1302, plain, (~ spl24_133 | spl24_95), inference(avatar_split_clause, [], [f499, f1082, f1298])).
fof(f499, plain, ((e4 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f274])).
fof(f1301, plain, (~ spl24_133 | ~ spl24_20), inference(avatar_split_clause, [], [f500, f767, f1298])).
fof(f500, plain, (~ (e4 = op(e4, e1)) | ~ sP21), inference(cnf_transformation, [], [f274])).
fof(f1296, plain, (~ spl24_132 | spl24_3), inference(avatar_split_clause, [], [f495, f696, f1291])).
fof(f495, plain, ((e2 = op(e4, e4)) | ~ sP22), inference(cnf_transformation, [], [f273])).
fof(f273, plain, ((~ (e4 = op(e4, e2)) & (e4 = op(e2, e2)) & (e2 = op(e4, e4))) | ~ sP22), inference(nnf_transformation, [], [f269])).
fof(f1295, plain, (~ spl24_132 | spl24_65), inference(avatar_split_clause, [], [f496, f956, f1291])).
fof(f496, plain, ((e4 = op(e2, e2)) | ~ sP22), inference(cnf_transformation, [], [f273])).
fof(f1294, plain, (~ spl24_132 | ~ spl24_15), inference(avatar_split_clause, [], [f497, f746, f1291])).
fof(f497, plain, (~ (e4 = op(e4, e2)) | ~ sP22), inference(cnf_transformation, [], [f273])).
fof(f1289, plain, (~ spl24_131 | spl24_4), inference(avatar_split_clause, [], [f492, f700, f1284])).
fof(f492, plain, ((e3 = op(e4, e4)) | ~ sP23), inference(cnf_transformation, [], [f272])).
fof(f272, plain, ((~ (e4 = op(e4, e3)) & (e4 = op(e3, e3)) & (e3 = op(e4, e4))) | ~ sP23), inference(nnf_transformation, [], [f270])).
fof(f1288, plain, (~ spl24_131 | spl24_35), inference(avatar_split_clause, [], [f493, f830, f1284])).
fof(f493, plain, ((e4 = op(e3, e3)) | ~ sP23), inference(cnf_transformation, [], [f272])).
fof(f1287, plain, (~ spl24_131 | ~ spl24_10), inference(avatar_split_clause, [], [f494, f725, f1284])).
fof(f494, plain, (~ (e4 = op(e4, e3)) | ~ sP23), inference(cnf_transformation, [], [f272])).
fof(f1280, plain, (spl24_122 | spl24_117 | spl24_112 | spl24_107 | spl24_102), inference(avatar_split_clause, [], [f334, f1112, f1133, f1154, f1175, f1196])).
fof(f334, plain, ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)) & ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)) & ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax3)).
fof(f1279, plain, (spl24_122 | spl24_97 | spl24_72 | spl24_47 | spl24_22), inference(avatar_split_clause, [], [f335, f776, f881, f986, f1091, f1196])).
fof(f335, plain, ((e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f3])).
fof(f1278, plain, (spl24_123 | spl24_118 | spl24_113 | spl24_108 | spl24_103), inference(avatar_split_clause, [], [f336, f1116, f1137, f1158, f1179, f1200])).
fof(f336, plain, ((e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f3])).
fof(f1277, plain, (spl24_123 | spl24_98 | spl24_73 | spl24_48 | spl24_23), inference(avatar_split_clause, [], [f337, f780, f885, f990, f1095, f1200])).
fof(f337, plain, ((e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f3])).
fof(f1276, plain, (spl24_124 | spl24_119 | spl24_114 | spl24_109 | spl24_104), inference(avatar_split_clause, [], [f338, f1120, f1141, f1162, f1183, f1204])).
fof(f338, plain, ((e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f3])).
fof(f1275, plain, (spl24_124 | spl24_99 | spl24_74 | spl24_49 | spl24_24), inference(avatar_split_clause, [], [f339, f784, f889, f994, f1099, f1204])).
fof(f339, plain, ((e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f3])).
fof(f1274, plain, (spl24_125 | spl24_120 | spl24_115 | spl24_110 | spl24_105), inference(avatar_split_clause, [], [f340, f1124, f1145, f1166, f1187, f1208])).
fof(f340, plain, ((e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f3])).
fof(f1273, plain, (spl24_125 | spl24_100 | spl24_75 | spl24_50 | spl24_25), inference(avatar_split_clause, [], [f341, f788, f893, f998, f1103, f1208])).
fof(f341, plain, ((e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (op(e0, e0) = e4)), inference(cnf_transformation, [], [f3])).
fof(f1272, plain, (spl24_96 | spl24_91 | spl24_86 | spl24_81 | spl24_76), inference(avatar_split_clause, [], [f342, f1003, f1024, f1045, f1066, f1087])).
fof(f342, plain, ((e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1271, plain, (spl24_116 | spl24_91 | spl24_66 | spl24_41 | spl24_16), inference(avatar_split_clause, [], [f343, f751, f856, f961, f1066, f1171])).
fof(f343, plain, ((e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1268, plain, (spl24_98 | spl24_93 | spl24_88 | spl24_83 | spl24_78), inference(avatar_split_clause, [], [f346, f1011, f1032, f1053, f1074, f1095])).
fof(f346, plain, ((e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1267, plain, (spl24_118 | spl24_93 | spl24_68 | spl24_43 | spl24_18), inference(avatar_split_clause, [], [f347, f759, f864, f969, f1074, f1179])).
fof(f347, plain, ((e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1266, plain, (spl24_99 | spl24_94 | spl24_89 | spl24_84 | spl24_79), inference(avatar_split_clause, [], [f348, f1015, f1036, f1057, f1078, f1099])).
fof(f348, plain, ((e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1265, plain, (spl24_119 | spl24_94 | spl24_69 | spl24_44 | spl24_19), inference(avatar_split_clause, [], [f349, f763, f868, f973, f1078, f1183])).
fof(f349, plain, ((e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1264, plain, (spl24_100 | spl24_95 | spl24_90 | spl24_85 | spl24_80), inference(avatar_split_clause, [], [f350, f1019, f1040, f1061, f1082, f1103])).
fof(f350, plain, ((e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))), inference(cnf_transformation, [], [f3])).
fof(f1263, plain, (spl24_120 | spl24_95 | spl24_70 | spl24_45 | spl24_20), inference(avatar_split_clause, [], [f351, f767, f872, f977, f1082, f1187])).
fof(f351, plain, ((e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))), inference(cnf_transformation, [], [f3])).
fof(f1262, plain, (spl24_71 | spl24_66 | spl24_61 | spl24_56 | spl24_51), inference(avatar_split_clause, [], [f352, f898, f919, f940, f961, f982])).
fof(f352, plain, ((e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1261, plain, (spl24_111 | spl24_86 | spl24_61 | spl24_36 | spl24_11), inference(avatar_split_clause, [], [f353, f730, f835, f940, f1045, f1150])).
fof(f353, plain, ((e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1260, plain, (spl24_72 | spl24_67 | spl24_62 | spl24_57 | spl24_52), inference(avatar_split_clause, [], [f354, f902, f923, f944, f965, f986])).
fof(f354, plain, ((e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1259, plain, (spl24_112 | spl24_87 | spl24_62 | spl24_37 | spl24_12), inference(avatar_split_clause, [], [f355, f734, f839, f944, f1049, f1154])).
fof(f355, plain, ((e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1256, plain, (spl24_74 | spl24_69 | spl24_64 | spl24_59 | spl24_54), inference(avatar_split_clause, [], [f358, f910, f931, f952, f973, f994])).
fof(f358, plain, ((e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1255, plain, (spl24_114 | spl24_89 | spl24_64 | spl24_39 | spl24_14), inference(avatar_split_clause, [], [f359, f742, f847, f952, f1057, f1162])).
fof(f359, plain, ((e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1254, plain, (spl24_75 | spl24_70 | spl24_65 | spl24_60 | spl24_55), inference(avatar_split_clause, [], [f360, f914, f935, f956, f977, f998])).
fof(f360, plain, ((e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))), inference(cnf_transformation, [], [f3])).
fof(f1253, plain, (spl24_115 | spl24_90 | spl24_65 | spl24_40 | spl24_15), inference(avatar_split_clause, [], [f361, f746, f851, f956, f1061, f1166])).
fof(f361, plain, ((e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))), inference(cnf_transformation, [], [f3])).
fof(f1252, plain, (spl24_46 | spl24_41 | spl24_36 | spl24_31 | spl24_26), inference(avatar_split_clause, [], [f362, f793, f814, f835, f856, f877])).
fof(f362, plain, ((e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1251, plain, (spl24_106 | spl24_81 | spl24_56 | spl24_31 | spl24_6), inference(avatar_split_clause, [], [f363, f709, f814, f919, f1024, f1129])).
fof(f363, plain, ((e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1250, plain, (spl24_47 | spl24_42 | spl24_37 | spl24_32 | spl24_27), inference(avatar_split_clause, [], [f364, f797, f818, f839, f860, f881])).
fof(f364, plain, ((e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1249, plain, (spl24_107 | spl24_82 | spl24_57 | spl24_32 | spl24_7), inference(avatar_split_clause, [], [f365, f713, f818, f923, f1028, f1133])).
fof(f365, plain, ((e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1248, plain, (spl24_48 | spl24_43 | spl24_38 | spl24_33 | spl24_28), inference(avatar_split_clause, [], [f366, f801, f822, f843, f864, f885])).
fof(f366, plain, ((e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1247, plain, (spl24_108 | spl24_83 | spl24_58 | spl24_33 | spl24_8), inference(avatar_split_clause, [], [f367, f717, f822, f927, f1032, f1137])).
fof(f367, plain, ((e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1244, plain, (spl24_50 | spl24_45 | spl24_40 | spl24_35 | spl24_30), inference(avatar_split_clause, [], [f370, f809, f830, f851, f872, f893])).
fof(f370, plain, ((e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))), inference(cnf_transformation, [], [f3])).
fof(f1243, plain, (spl24_110 | spl24_85 | spl24_60 | spl24_35 | spl24_10), inference(avatar_split_clause, [], [f371, f725, f830, f935, f1040, f1145])).
fof(f371, plain, ((e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))), inference(cnf_transformation, [], [f3])).
fof(f1242, plain, (spl24_21 | spl24_16 | spl24_11 | spl24_6 | spl24_1), inference(avatar_split_clause, [], [f372, f688, f709, f730, f751, f772])).
fof(f372, plain, ((e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1241, plain, (spl24_101 | spl24_76 | spl24_51 | spl24_26 | spl24_1), inference(avatar_split_clause, [], [f373, f688, f793, f898, f1003, f1108])).
fof(f373, plain, ((e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1240, plain, (spl24_22 | spl24_17 | spl24_12 | spl24_7 | spl24_2), inference(avatar_split_clause, [], [f374, f692, f713, f734, f755, f776])).
fof(f374, plain, ((e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1239, plain, (spl24_102 | spl24_77 | spl24_52 | spl24_27 | spl24_2), inference(avatar_split_clause, [], [f375, f692, f797, f902, f1007, f1112])).
fof(f375, plain, ((e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1238, plain, (spl24_23 | spl24_18 | spl24_13 | spl24_8 | spl24_3), inference(avatar_split_clause, [], [f376, f696, f717, f738, f759, f780])).
fof(f376, plain, ((e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1237, plain, (spl24_103 | spl24_78 | spl24_53 | spl24_28 | spl24_3), inference(avatar_split_clause, [], [f377, f696, f801, f906, f1011, f1116])).
fof(f377, plain, ((e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1236, plain, (spl24_24 | spl24_19 | spl24_14 | spl24_9 | spl24_4), inference(avatar_split_clause, [], [f378, f700, f721, f742, f763, f784])).
fof(f378, plain, ((e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))), inference(cnf_transformation, [], [f3])).
fof(f1235, plain, (spl24_104 | spl24_79 | spl24_54 | spl24_29 | spl24_4), inference(avatar_split_clause, [], [f379, f700, f805, f910, f1015, f1120])).
fof(f379, plain, ((e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))), inference(cnf_transformation, [], [f3])).
fof(f1232, plain, (spl24_126 | spl24_127 | spl24_128 | spl24_129 | spl24_130), inference(avatar_split_clause, [], [f331, f1229, f1225, f1221, f1217, f1213])).
fof(f331, plain, ((e4 = unit) | (e3 = unit) | (e2 = unit) | (e1 = unit) | (e0 = unit)), inference(cnf_transformation, [], [f2])).
fof(f1211, plain, (spl24_121 | spl24_122 | spl24_123 | spl24_124 | spl24_125), inference(avatar_split_clause, [], [f296, f1208, f1204, f1200, f1196, f1192])).
fof(f296, plain, ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e4) | (op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG100+1.p', ax1)).
fof(f1190, plain, (spl24_116 | spl24_117 | spl24_118 | spl24_119 | spl24_120), inference(avatar_split_clause, [], [f297, f1187, f1183, f1179, f1175, f1171])).
fof(f297, plain, ((e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f1169, plain, (spl24_111 | spl24_112 | spl24_113 | spl24_114 | spl24_115), inference(avatar_split_clause, [], [f298, f1166, f1162, f1158, f1154, f1150])).
fof(f298, plain, ((e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f1148, plain, (spl24_106 | spl24_107 | spl24_108 | spl24_109 | spl24_110), inference(avatar_split_clause, [], [f299, f1145, f1141, f1137, f1133, f1129])).
fof(f299, plain, ((e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f1127, plain, (spl24_101 | spl24_102 | spl24_103 | spl24_104 | spl24_105), inference(avatar_split_clause, [], [f300, f1124, f1120, f1116, f1112, f1108])).
fof(f300, plain, ((e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))), inference(cnf_transformation, [], [f1])).
fof(f1106, plain, (spl24_96 | spl24_97 | spl24_98 | spl24_99 | spl24_100), inference(avatar_split_clause, [], [f301, f1103, f1099, f1095, f1091, f1087])).
fof(f301, plain, ((e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f1085, plain, (spl24_91 | spl24_92 | spl24_93 | spl24_94 | spl24_95), inference(avatar_split_clause, [], [f302, f1082, f1078, f1074, f1070, f1066])).
fof(f302, plain, ((e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f1064, plain, (spl24_86 | spl24_87 | spl24_88 | spl24_89 | spl24_90), inference(avatar_split_clause, [], [f303, f1061, f1057, f1053, f1049, f1045])).
fof(f303, plain, ((e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f1043, plain, (spl24_81 | spl24_82 | spl24_83 | spl24_84 | spl24_85), inference(avatar_split_clause, [], [f304, f1040, f1036, f1032, f1028, f1024])).
fof(f304, plain, ((e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f1022, plain, (spl24_76 | spl24_77 | spl24_78 | spl24_79 | spl24_80), inference(avatar_split_clause, [], [f305, f1019, f1015, f1011, f1007, f1003])).
fof(f305, plain, ((e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))), inference(cnf_transformation, [], [f1])).
fof(f1001, plain, (spl24_71 | spl24_72 | spl24_73 | spl24_74 | spl24_75), inference(avatar_split_clause, [], [f306, f998, f994, f990, f986, f982])).
fof(f306, plain, ((e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f980, plain, (spl24_66 | spl24_67 | spl24_68 | spl24_69 | spl24_70), inference(avatar_split_clause, [], [f307, f977, f973, f969, f965, f961])).
fof(f307, plain, ((e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f959, plain, (spl24_61 | spl24_62 | spl24_63 | spl24_64 | spl24_65), inference(avatar_split_clause, [], [f308, f956, f952, f948, f944, f940])).
fof(f308, plain, ((e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f938, plain, (spl24_56 | spl24_57 | spl24_58 | spl24_59 | spl24_60), inference(avatar_split_clause, [], [f309, f935, f931, f927, f923, f919])).
fof(f309, plain, ((e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f917, plain, (spl24_51 | spl24_52 | spl24_53 | spl24_54 | spl24_55), inference(avatar_split_clause, [], [f310, f914, f910, f906, f902, f898])).
fof(f310, plain, ((e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))), inference(cnf_transformation, [], [f1])).
fof(f896, plain, (spl24_46 | spl24_47 | spl24_48 | spl24_49 | spl24_50), inference(avatar_split_clause, [], [f311, f893, f889, f885, f881, f877])).
fof(f311, plain, ((e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f875, plain, (spl24_41 | spl24_42 | spl24_43 | spl24_44 | spl24_45), inference(avatar_split_clause, [], [f312, f872, f868, f864, f860, f856])).
fof(f312, plain, ((e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f854, plain, (spl24_36 | spl24_37 | spl24_38 | spl24_39 | spl24_40), inference(avatar_split_clause, [], [f313, f851, f847, f843, f839, f835])).
fof(f313, plain, ((e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f812, plain, (spl24_26 | spl24_27 | spl24_28 | spl24_29 | spl24_30), inference(avatar_split_clause, [], [f315, f809, f805, f801, f797, f793])).
fof(f315, plain, ((e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))), inference(cnf_transformation, [], [f1])).
fof(f791, plain, (spl24_21 | spl24_22 | spl24_23 | spl24_24 | spl24_25), inference(avatar_split_clause, [], [f316, f788, f784, f780, f776, f772])).
fof(f316, plain, ((e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))), inference(cnf_transformation, [], [f1])).
fof(f770, plain, (spl24_16 | spl24_17 | spl24_18 | spl24_19 | spl24_20), inference(avatar_split_clause, [], [f317, f767, f763, f759, f755, f751])).
fof(f317, plain, ((e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))), inference(cnf_transformation, [], [f1])).
fof(f749, plain, (spl24_11 | spl24_12 | spl24_13 | spl24_14 | spl24_15), inference(avatar_split_clause, [], [f318, f746, f742, f738, f734, f730])).
fof(f318, plain, ((e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))), inference(cnf_transformation, [], [f1])).
fof(f728, plain, (spl24_6 | spl24_7 | spl24_8 | spl24_9 | spl24_10), inference(avatar_split_clause, [], [f319, f725, f721, f717, f713, f709])).
fof(f319, plain, ((e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))), inference(cnf_transformation, [], [f1])).
fof(f707, plain, (spl24_1 | spl24_2 | spl24_3 | spl24_4 | spl24_5), inference(avatar_split_clause, [], [f320, f704, f700, f696, f692, f688])).
fof(f320, plain, ((e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))), inference(cnf_transformation, [], [f1])).