fof(f15108, plain, $false, inference(avatar_sat_refutation, [], [f1699, f1741, f1825, f1827, f1834, f1838, f1847, f1851, f1867, f2022, f2274, f2400, f2402, f2421, f2426, f2427, f2436, f2451, f2453, f2454, f2456, f3801, f3825, f3845, f3865, f7101, f8407, f8468, f8469, f8470, f8475, f8504, f8540, f8566, f8624, f8662, f8726, f8864, f8909, f8917, f8934, f8946, f8950, f8951, f8964, f9013, f9041, f9057, f9086, f9088, f9121, f9124, f9126, f9162, f9170, f9213, f9215, f9226, f9262, f9296, f9300, f9306, f9314, f9328, f9376, f9377, f9378, f9379, f9390, f9391, f9406, f9439, f9460, f9484, f9494, f9534, f9708, f9778, f9786, f9794, f9796, f9820, f9822, f9861, f9872, f9945, f9949, f9965, f10006, f10013, f10019, f10044, f10067, f10095, f10111, f10112, f10130, f10139, f10190, f10228, f10236, f10239, f10241, f10272, f10325, f10365, f10368, f10370, f10451, f10468, f10486, f10549, f10550, f10602, f10616, f10635, f10645, f10654, f10671, f10713, f10784, f10790, f10791, f10809, f10810, f10820, f10832, f10845, f10858, f10881, f10953, f11001, f11024, f11068, f11079, f11201, f11278, f11344, f11400, f11449, f11458, f11476, f11572, f11601, f11627, f11629, f11639, f11668, f11722, f11728, f11730, f11734, f11773, f11893, f12060, f12239, f12426, f12598, f12765, f12939, f13112, f13272, f13416, f13552, f13689, f13815, f13946, f14074, f14197, f14318, f14447, f14579, f14719, f14867, f15013])).
fof(f15013, plain, (~ spl80_121 | ~ spl80_246 | spl80_952), inference(avatar_contradiction_clause, [], [f15012])).
fof(f15012, plain, ($false | (~ spl80_121 | ~ spl80_246 | spl80_952)), inference(subsumption_resolution, [], [f15011, f2383])).
fof(f2383, plain, ((e20 = op2(e20, e20)) | ~ spl80_246), inference(avatar_component_clause, [], [f2381])).
fof(f2381, plain, (spl80_246 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_246])])).
fof(f15011, plain, (~ (e20 = op2(e20, e20)) | (~ spl80_121 | spl80_952)), inference(forward_demodulation, [], [f15010, f2546])).
fof(f2546, plain, (e20 = h7(e10)), inference(forward_demodulation, [], [f2545, f2541])).
fof(f2541, plain, (e20 = op2(h7(e14), e23)), inference(backward_demodulation, [], [f698, f735])).
fof(f735, plain, (op2(e23, e21) = h7(e14)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ((op2(e23, e21) = h7(e14)) & (op2(e23, op2(op2(e23, e21), e23)) = h7(e12)) & (op2(op2(e23, e21), e23) = h7(e10)) & (e23 = h7(e13)) & (e21 = h7(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax22)).
fof(f698, plain, (e20 = op2(op2(e23, e21), e23)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((e24 = op2(e23, e21)) & (e22 = op2(e23, op2(op2(e23, e21), e23))) & (e20 = op2(op2(e23, e21), e23))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax15)).
fof(f2545, plain, (h7(e10) = op2(h7(e14), e23)), inference(forward_demodulation, [], [f733, f735])).
fof(f733, plain, (op2(op2(e23, e21), e23) = h7(e10)), inference(cnf_transformation, [], [f22])).
fof(f15010, plain, (~ (op2(e20, e20) = h7(e10)) | (~ spl80_121 | spl80_952)), inference(forward_demodulation, [], [f7100, f1808])).
fof(f1808, plain, ((e10 = op1(e10, e10)) | ~ spl80_121), inference(avatar_component_clause, [], [f1806])).
fof(f1806, plain, (spl80_121 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_121])])).
fof(f7100, plain, (~ (op2(e20, e20) = h7(op1(e10, e10))) | spl80_952), inference(avatar_component_clause, [], [f7098])).
fof(f7098, plain, (spl80_952 <=> (op2(e20, e20) = h7(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_952])])).
fof(f14867, plain, (~ spl80_119 | ~ spl80_488 | spl80_951), inference(avatar_contradiction_clause, [], [f14866])).
fof(f14866, plain, ($false | (~ spl80_119 | ~ spl80_488 | spl80_951)), inference(subsumption_resolution, [], [f14865, f3958])).
fof(f3958, plain, ((e23 = h5(e14)) | ~ spl80_488), inference(avatar_component_clause, [], [f3957])).
fof(f3957, plain, (spl80_488 <=> (e23 = h5(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_488])])).
fof(f14865, plain, (~ (e23 = h5(e14)) | (~ spl80_119 | spl80_951)), inference(forward_demodulation, [], [f14864, f732])).
fof(f732, plain, (e23 = h7(e13)), inference(cnf_transformation, [], [f22])).
fof(f14864, plain, (~ (h5(e14) = h7(e13)) | (~ spl80_119 | spl80_951)), inference(forward_demodulation, [], [f7096, f1799])).
fof(f1799, plain, ((e13 = op1(e10, e11)) | ~ spl80_119), inference(avatar_component_clause, [], [f1797])).
fof(f1797, plain, (spl80_119 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_119])])).
fof(f7096, plain, (~ (h5(e14) = h7(op1(e10, e11))) | spl80_951), inference(avatar_component_clause, [], [f7094])).
fof(f7094, plain, (spl80_951 <=> (h5(e14) = h7(op1(e10, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_951])])).
fof(f14719, plain, (~ spl80_107 | ~ spl80_372 | spl80_949), inference(avatar_contradiction_clause, [], [f14718])).
fof(f14718, plain, ($false | (~ spl80_107 | ~ spl80_372 | spl80_949)), inference(subsumption_resolution, [], [f14717, f731])).
fof(f731, plain, (e21 = h7(e11)), inference(cnf_transformation, [], [f22])).
fof(f14717, plain, (~ (e21 = h7(e11)) | (~ spl80_107 | ~ spl80_372 | spl80_949)), inference(forward_demodulation, [], [f14716, f3326])).
fof(f3326, plain, ((e21 = h13(e14)) | ~ spl80_372), inference(avatar_component_clause, [], [f3325])).
fof(f3325, plain, (spl80_372 <=> (e21 = h13(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_372])])).
fof(f14716, plain, (~ (h7(e11) = h13(e14)) | (~ spl80_107 | spl80_949)), inference(forward_demodulation, [], [f7088, f1749])).
fof(f1749, plain, ((e11 = op1(e10, e13)) | ~ spl80_107), inference(avatar_component_clause, [], [f1747])).
fof(f1747, plain, (spl80_107 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_107])])).
fof(f7088, plain, (~ (h13(e14) = h7(op1(e10, e13))) | spl80_949), inference(avatar_component_clause, [], [f7086])).
fof(f7086, plain, (spl80_949 <=> (h13(e14) = h7(op1(e10, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_949])])).
fof(f14579, plain, (~ spl80_103 | ~ spl80_228 | ~ spl80_464 | ~ spl80_928 | spl80_948), inference(avatar_contradiction_clause, [], [f14578])).
fof(f14578, plain, ($false | (~ spl80_103 | ~ spl80_228 | ~ spl80_464 | ~ spl80_928 | spl80_948)), inference(subsumption_resolution, [], [f14577, f2307])).
fof(f2307, plain, ((e22 = op2(e20, e24)) | ~ spl80_228), inference(avatar_component_clause, [], [f2305])).
fof(f2305, plain, (spl80_228 <=> (e22 = op2(e20, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_228])])).
fof(f14577, plain, (~ (e22 = op2(e20, e24)) | (~ spl80_103 | ~ spl80_464 | ~ spl80_928 | spl80_948)), inference(forward_demodulation, [], [f14576, f9385])).
fof(f9385, plain, ((e22 = h7(e12)) | ~ spl80_464), inference(backward_demodulation, [], [f2544, f3823])).
fof(f3823, plain, ((e22 = h3(e14)) | ~ spl80_464), inference(avatar_component_clause, [], [f3822])).
fof(f3822, plain, (spl80_464 <=> (e22 = h3(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_464])])).
fof(f2544, plain, (h3(e14) = h7(e12)), inference(forward_demodulation, [], [f2543, f715])).
fof(f715, plain, (op2(e23, e20) = h3(e14)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ((op2(e23, e20) = h3(e14)) & (h3(e12) = op2(e23, op2(op2(e23, e20), e23))) & (h3(e10) = op2(op2(e23, e20), e23)) & (e23 = h3(e13)) & (e20 = h3(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax18)).
fof(f2543, plain, (op2(e23, e20) = h7(e12)), inference(forward_demodulation, [], [f2542, f2541])).
fof(f2542, plain, (h7(e12) = op2(e23, op2(h7(e14), e23))), inference(forward_demodulation, [], [f734, f735])).
fof(f734, plain, (op2(e23, op2(op2(e23, e21), e23)) = h7(e12)), inference(cnf_transformation, [], [f22])).
fof(f14576, plain, (~ (op2(e20, e24) = h7(e12)) | (~ spl80_103 | ~ spl80_928 | spl80_948)), inference(forward_demodulation, [], [f14575, f1732])).
fof(f1732, plain, ((e12 = op1(e10, e14)) | ~ spl80_103), inference(avatar_component_clause, [], [f1730])).
fof(f1730, plain, (spl80_103 <=> (e12 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_103])])).
fof(f14575, plain, (~ (op2(e20, e24) = h7(op1(e10, e14))) | (~ spl80_928 | spl80_948)), inference(forward_demodulation, [], [f7084, f7003])).
fof(f7003, plain, ((e24 = h7(e14)) | ~ spl80_928), inference(avatar_component_clause, [], [f7002])).
fof(f7002, plain, (spl80_928 <=> (e24 = h7(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_928])])).
fof(f7084, plain, (~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | spl80_948), inference(avatar_component_clause, [], [f7082])).
fof(f7082, plain, (spl80_948 <=> (h7(op1(e10, e14)) = op2(e20, h7(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_948])])).
fof(f14447, plain, (~ spl80_100 | ~ spl80_928 | spl80_947 | ~ spl80_1102), inference(avatar_contradiction_clause, [], [f14446])).
fof(f14446, plain, ($false | (~ spl80_100 | ~ spl80_928 | spl80_947 | ~ spl80_1102)), inference(subsumption_resolution, [], [f14445, f8168])).
fof(f8168, plain, ((e24 = h1(e14)) | ~ spl80_1102), inference(avatar_component_clause, [], [f8167])).
fof(f8167, plain, (spl80_1102 <=> (e24 = h1(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_1102])])).
fof(f14445, plain, (~ (e24 = h1(e14)) | (~ spl80_100 | ~ spl80_928 | spl80_947)), inference(forward_demodulation, [], [f14444, f7003])).
fof(f14444, plain, (~ (h1(e14) = h7(e14)) | (~ spl80_100 | spl80_947)), inference(forward_demodulation, [], [f7080, f1719])).
fof(f1719, plain, ((e14 = op1(e11, e10)) | ~ spl80_100), inference(avatar_component_clause, [], [f1717])).
fof(f1717, plain, (spl80_100 <=> (e14 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_100])])).
fof(f7080, plain, (~ (h1(e14) = h7(op1(e11, e10))) | spl80_947), inference(avatar_component_clause, [], [f7078])).
fof(f7078, plain, (spl80_947 <=> (h1(e14) = h7(op1(e11, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_947])])).
fof(f14318, plain, (~ spl80_92 | ~ spl80_217 | spl80_946), inference(avatar_contradiction_clause, [], [f14317])).
fof(f14317, plain, ($false | (~ spl80_92 | ~ spl80_217 | spl80_946)), inference(subsumption_resolution, [], [f14316, f2261])).
fof(f2261, plain, ((e21 = op2(e21, e21)) | ~ spl80_217), inference(avatar_component_clause, [], [f2259])).
fof(f2259, plain, (spl80_217 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_217])])).
fof(f14316, plain, (~ (e21 = op2(e21, e21)) | (~ spl80_92 | spl80_946)), inference(forward_demodulation, [], [f14315, f731])).
fof(f14315, plain, (~ (op2(e21, e21) = h7(e11)) | (~ spl80_92 | spl80_946)), inference(forward_demodulation, [], [f7076, f1686])).
fof(f1686, plain, ((e11 = op1(e11, e11)) | ~ spl80_92), inference(avatar_component_clause, [], [f1684])).
fof(f1684, plain, (spl80_92 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_92])])).
fof(f7076, plain, (~ (op2(e21, e21) = h7(op1(e11, e11))) | spl80_946), inference(avatar_component_clause, [], [f7074])).
fof(f7074, plain, (spl80_946 <=> (op2(e21, e21) = h7(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_946])])).
fof(f14197, plain, (~ spl80_89 | ~ spl80_214 | ~ spl80_464 | spl80_945), inference(avatar_contradiction_clause, [], [f14196])).
fof(f14196, plain, ($false | (~ spl80_89 | ~ spl80_214 | ~ spl80_464 | spl80_945)), inference(subsumption_resolution, [], [f14195, f2248])).
fof(f2248, plain, ((e23 = op2(e21, e22)) | ~ spl80_214), inference(avatar_component_clause, [], [f2246])).
fof(f2246, plain, (spl80_214 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_214])])).
fof(f14195, plain, (~ (e23 = op2(e21, e22)) | (~ spl80_89 | ~ spl80_464 | spl80_945)), inference(forward_demodulation, [], [f14194, f732])).
fof(f14194, plain, (~ (op2(e21, e22) = h7(e13)) | (~ spl80_89 | ~ spl80_464 | spl80_945)), inference(forward_demodulation, [], [f14193, f1673])).
fof(f1673, plain, ((e13 = op1(e11, e12)) | ~ spl80_89), inference(avatar_component_clause, [], [f1671])).
fof(f1671, plain, (spl80_89 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_89])])).
fof(f14193, plain, (~ (op2(e21, e22) = h7(op1(e11, e12))) | (~ spl80_464 | spl80_945)), inference(forward_demodulation, [], [f7072, f3823])).
fof(f7072, plain, (~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | spl80_945), inference(avatar_component_clause, [], [f7070])).
fof(f7070, plain, (spl80_945 <=> (h7(op1(e11, e12)) = op2(e21, h3(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_945])])).
fof(f14074, plain, (~ spl80_76 | ~ spl80_201 | ~ spl80_928 | spl80_943), inference(avatar_contradiction_clause, [], [f14073])).
fof(f14073, plain, ($false | (~ spl80_76 | ~ spl80_201 | ~ spl80_928 | spl80_943)), inference(subsumption_resolution, [], [f14072, f2194])).
fof(f2194, plain, ((e20 = op2(e21, e24)) | ~ spl80_201), inference(avatar_component_clause, [], [f2192])).
fof(f2192, plain, (spl80_201 <=> (e20 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_201])])).
fof(f14072, plain, (~ (e20 = op2(e21, e24)) | (~ spl80_76 | ~ spl80_928 | spl80_943)), inference(forward_demodulation, [], [f14071, f2546])).
fof(f14071, plain, (~ (op2(e21, e24) = h7(e10)) | (~ spl80_76 | ~ spl80_928 | spl80_943)), inference(forward_demodulation, [], [f14070, f1619])).
fof(f1619, plain, ((e10 = op1(e11, e14)) | ~ spl80_76), inference(avatar_component_clause, [], [f1617])).
fof(f1617, plain, (spl80_76 <=> (e10 = op1(e11, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_76])])).
fof(f14070, plain, (~ (op2(e21, e24) = h7(op1(e11, e14))) | (~ spl80_928 | spl80_943)), inference(forward_demodulation, [], [f7064, f7003])).
fof(f7064, plain, (~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | spl80_943), inference(avatar_component_clause, [], [f7062])).
fof(f7062, plain, (spl80_943 <=> (h7(op1(e11, e14)) = op2(e21, h7(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_943])])).
fof(f13946, plain, (~ spl80_72 | ~ spl80_197 | ~ spl80_464 | spl80_942), inference(avatar_contradiction_clause, [], [f13945])).
fof(f13945, plain, ($false | (~ spl80_72 | ~ spl80_197 | ~ spl80_464 | spl80_942)), inference(subsumption_resolution, [], [f13944, f2177])).
fof(f2177, plain, ((e21 = op2(e22, e20)) | ~ spl80_197), inference(avatar_component_clause, [], [f2175])).
fof(f2175, plain, (spl80_197 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_197])])).
fof(f13944, plain, (~ (e21 = op2(e22, e20)) | (~ spl80_72 | ~ spl80_464 | spl80_942)), inference(forward_demodulation, [], [f13943, f731])).
fof(f13943, plain, (~ (op2(e22, e20) = h7(e11)) | (~ spl80_72 | ~ spl80_464 | spl80_942)), inference(forward_demodulation, [], [f13942, f1602])).
fof(f1602, plain, ((e11 = op1(e12, e10)) | ~ spl80_72), inference(avatar_component_clause, [], [f1600])).
fof(f1600, plain, (spl80_72 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_72])])).
fof(f13942, plain, (~ (op2(e22, e20) = h7(op1(e12, e10))) | (~ spl80_464 | spl80_942)), inference(forward_demodulation, [], [f7060, f3823])).
fof(f7060, plain, (~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | spl80_942), inference(avatar_component_clause, [], [f7058])).
fof(f7058, plain, (spl80_942 <=> (h7(op1(e12, e10)) = op2(h3(e14), e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_942])])).
fof(f13815, plain, (~ spl80_66 | ~ spl80_191 | ~ spl80_464 | spl80_941), inference(avatar_contradiction_clause, [], [f13814])).
fof(f13814, plain, ($false | (~ spl80_66 | ~ spl80_191 | ~ spl80_464 | spl80_941)), inference(subsumption_resolution, [], [f13813, f2152])).
fof(f2152, plain, ((e20 = op2(e22, e21)) | ~ spl80_191), inference(avatar_component_clause, [], [f2150])).
fof(f2150, plain, (spl80_191 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_191])])).
fof(f13813, plain, (~ (e20 = op2(e22, e21)) | (~ spl80_66 | ~ spl80_464 | spl80_941)), inference(forward_demodulation, [], [f13812, f2546])).
fof(f13812, plain, (~ (op2(e22, e21) = h7(e10)) | (~ spl80_66 | ~ spl80_464 | spl80_941)), inference(forward_demodulation, [], [f13811, f1577])).
fof(f1577, plain, ((e10 = op1(e12, e11)) | ~ spl80_66), inference(avatar_component_clause, [], [f1575])).
fof(f1575, plain, (spl80_66 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_66])])).
fof(f13811, plain, (~ (op2(e22, e21) = h7(op1(e12, e11))) | (~ spl80_464 | spl80_941)), inference(forward_demodulation, [], [f7056, f3823])).
fof(f7056, plain, (~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | spl80_941), inference(avatar_component_clause, [], [f7054])).
fof(f7054, plain, (spl80_941 <=> (h7(op1(e12, e11)) = op2(h3(e14), e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_941])])).
fof(f13689, plain, (~ spl80_63 | ~ spl80_188 | ~ spl80_464 | spl80_940), inference(avatar_contradiction_clause, [], [f13688])).
fof(f13688, plain, ($false | (~ spl80_63 | ~ spl80_188 | ~ spl80_464 | spl80_940)), inference(subsumption_resolution, [], [f13687, f2139])).
fof(f2139, plain, ((e22 = op2(e22, e22)) | ~ spl80_188), inference(avatar_component_clause, [], [f2137])).
fof(f2137, plain, (spl80_188 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_188])])).
fof(f13687, plain, (~ (e22 = op2(e22, e22)) | (~ spl80_63 | ~ spl80_464 | spl80_940)), inference(forward_demodulation, [], [f13686, f9385])).
fof(f13686, plain, (~ (op2(e22, e22) = h7(e12)) | (~ spl80_63 | ~ spl80_464 | spl80_940)), inference(forward_demodulation, [], [f13685, f3823])).
fof(f13685, plain, (~ (h7(e12) = op2(h3(e14), h3(e14))) | (~ spl80_63 | spl80_940)), inference(forward_demodulation, [], [f7052, f1564])).
fof(f1564, plain, ((e12 = op1(e12, e12)) | ~ spl80_63), inference(avatar_component_clause, [], [f1562])).
fof(f1562, plain, (spl80_63 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_63])])).
fof(f7052, plain, (~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | spl80_940), inference(avatar_component_clause, [], [f7050])).
fof(f7050, plain, (spl80_940 <=> (op2(h3(e14), h3(e14)) = h7(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl80_940])])).
fof(f13552, plain, (~ spl80_60 | ~ spl80_928 | spl80_939 | ~ spl80_1064), inference(avatar_contradiction_clause, [], [f13551])).
fof(f13551, plain, ($false | (~ spl80_60 | ~ spl80_928 | spl80_939 | ~ spl80_1064)), inference(subsumption_resolution, [], [f13550, f7918])).
fof(f7918, plain, ((e24 = h3(e10)) | ~ spl80_1064), inference(avatar_component_clause, [], [f7917])).
fof(f7917, plain, (spl80_1064 <=> (e24 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_1064])])).
fof(f13550, plain, (~ (e24 = h3(e10)) | (~ spl80_60 | ~ spl80_928 | spl80_939)), inference(forward_demodulation, [], [f13549, f7003])).
fof(f13549, plain, (~ (h3(e10) = h7(e14)) | (~ spl80_60 | spl80_939)), inference(forward_demodulation, [], [f7048, f1551])).
fof(f1551, plain, ((e14 = op1(e12, e13)) | ~ spl80_60), inference(avatar_component_clause, [], [f1549])).
fof(f1549, plain, (spl80_60 <=> (e14 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_60])])).
fof(f7048, plain, (~ (h3(e10) = h7(op1(e12, e13))) | spl80_939), inference(avatar_component_clause, [], [f7046])).
fof(f7046, plain, (spl80_939 <=> (h3(e10) = h7(op1(e12, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_939])])).
fof(f13416, plain, (~ spl80_54 | ~ spl80_179 | ~ spl80_464 | ~ spl80_928 | spl80_938), inference(avatar_contradiction_clause, [], [f13415])).
fof(f13415, plain, ($false | (~ spl80_54 | ~ spl80_179 | ~ spl80_464 | ~ spl80_928 | spl80_938)), inference(subsumption_resolution, [], [f13414, f2101])).
fof(f2101, plain, ((e23 = op2(e22, e24)) | ~ spl80_179), inference(avatar_component_clause, [], [f2099])).
fof(f2099, plain, (spl80_179 <=> (e23 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_179])])).
fof(f13414, plain, (~ (e23 = op2(e22, e24)) | (~ spl80_54 | ~ spl80_464 | ~ spl80_928 | spl80_938)), inference(forward_demodulation, [], [f13413, f732])).
fof(f13413, plain, (~ (op2(e22, e24) = h7(e13)) | (~ spl80_54 | ~ spl80_464 | ~ spl80_928 | spl80_938)), inference(forward_demodulation, [], [f13412, f1526])).
fof(f1526, plain, ((e13 = op1(e12, e14)) | ~ spl80_54), inference(avatar_component_clause, [], [f1524])).
fof(f1524, plain, (spl80_54 <=> (e13 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_54])])).
fof(f13412, plain, (~ (op2(e22, e24) = h7(op1(e12, e14))) | (~ spl80_464 | ~ spl80_928 | spl80_938)), inference(forward_demodulation, [], [f13411, f3823])).
fof(f13411, plain, (~ (h7(op1(e12, e14)) = op2(h3(e14), e24)) | (~ spl80_928 | spl80_938)), inference(forward_demodulation, [], [f7044, f7003])).
fof(f7044, plain, (~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | spl80_938), inference(avatar_component_clause, [], [f7042])).
fof(f7042, plain, (spl80_938 <=> (h7(op1(e12, e14)) = op2(h3(e14), h7(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_938])])).
fof(f13272, plain, (~ spl80_48 | ~ spl80_464 | spl80_937), inference(avatar_contradiction_clause, [], [f13271])).
fof(f13271, plain, ($false | (~ spl80_48 | ~ spl80_464 | spl80_937)), inference(subsumption_resolution, [], [f13270, f3823])).
fof(f13270, plain, (~ (e22 = h3(e14)) | (~ spl80_48 | ~ spl80_464 | spl80_937)), inference(forward_demodulation, [], [f13269, f9385])).
fof(f13269, plain, (~ (h3(e14) = h7(e12)) | (~ spl80_48 | spl80_937)), inference(forward_demodulation, [], [f7040, f1501])).
fof(f1501, plain, ((e12 = op1(e13, e10)) | ~ spl80_48), inference(avatar_component_clause, [], [f1499])).
fof(f1499, plain, (spl80_48 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_48])])).
fof(f7040, plain, (~ (h3(e14) = h7(op1(e13, e10))) | spl80_937), inference(avatar_component_clause, [], [f7038])).
fof(f7038, plain, (spl80_937 <=> (h3(e14) = h7(op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl80_937])])).
fof(f13112, plain, (~ spl80_83 | ~ spl80_352 | ~ spl80_464 | spl80_944), inference(avatar_contradiction_clause, [], [f13111])).
fof(f13111, plain, ($false | (~ spl80_83 | ~ spl80_352 | ~ spl80_464 | spl80_944)), inference(subsumption_resolution, [], [f13110, f9385])).
fof(f13110, plain, (~ (e22 = h7(e12)) | (~ spl80_83 | ~ spl80_352 | spl80_944)), inference(forward_demodulation, [], [f13109, f3219])).
fof(f3219, plain, ((e22 = h14(e14)) | ~ spl80_352), inference(avatar_component_clause, [], [f3218])).
fof(f3218, plain, (spl80_352 <=> (e22 = h14(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_352])])).
fof(f13109, plain, (~ (h7(e12) = h14(e14)) | (~ spl80_83 | spl80_944)), inference(forward_demodulation, [], [f7068, f1648])).
fof(f1648, plain, ((e12 = op1(e11, e13)) | ~ spl80_83), inference(avatar_component_clause, [], [f1646])).
fof(f1646, plain, (spl80_83 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_83])])).
fof(f7068, plain, (~ (h14(e14) = h7(op1(e11, e13))) | spl80_944), inference(avatar_component_clause, [], [f7066])).
fof(f7066, plain, (spl80_944 <=> (h14(e14) = h7(op1(e11, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_944])])).
fof(f12939, plain, (~ spl80_45 | spl80_936), inference(avatar_contradiction_clause, [], [f12938])).
fof(f12938, plain, ($false | (~ spl80_45 | spl80_936)), inference(trivial_inequality_removal, [], [f12937])).
fof(f12937, plain, (~ (h7(e14) = h7(e14)) | (~ spl80_45 | spl80_936)), inference(forward_demodulation, [], [f7036, f1488])).
fof(f1488, plain, ((e14 = op1(e13, e11)) | ~ spl80_45), inference(avatar_component_clause, [], [f1486])).
fof(f1486, plain, (spl80_45 <=> (e14 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_45])])).
fof(f7036, plain, (~ (h7(e14) = h7(op1(e13, e11))) | spl80_936), inference(avatar_component_clause, [], [f7034])).
fof(f7034, plain, (spl80_936 <=> (h7(e14) = h7(op1(e13, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl80_936])])).
fof(f12765, plain, (~ spl80_36 | spl80_935), inference(avatar_contradiction_clause, [], [f12764])).
fof(f12764, plain, ($false | (~ spl80_36 | spl80_935)), inference(subsumption_resolution, [], [f12763, f2546])).
fof(f12763, plain, (~ (e20 = h7(e10)) | (~ spl80_36 | spl80_935)), inference(forward_demodulation, [], [f7032, f1451])).
fof(f1451, plain, ((e10 = op1(e13, e12)) | ~ spl80_36), inference(avatar_component_clause, [], [f1449])).
fof(f1449, plain, (spl80_36 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_36])])).
fof(f7032, plain, (~ (e20 = h7(op1(e13, e12))) | spl80_935), inference(avatar_component_clause, [], [f7030])).
fof(f7030, plain, (spl80_935 <=> (e20 = h7(op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl80_935])])).
fof(f12598, plain, (~ spl80_34 | ~ spl80_159 | spl80_934), inference(avatar_contradiction_clause, [], [f12597])).
fof(f12597, plain, ($false | (~ spl80_34 | ~ spl80_159 | spl80_934)), inference(subsumption_resolution, [], [f12596, f2017])).
fof(f2017, plain, ((e23 = op2(e23, e23)) | ~ spl80_159), inference(avatar_component_clause, [], [f2015])).
fof(f2015, plain, (spl80_159 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_159])])).
fof(f12596, plain, (~ (e23 = op2(e23, e23)) | (~ spl80_34 | spl80_934)), inference(forward_demodulation, [], [f12595, f732])).
fof(f12595, plain, (~ (op2(e23, e23) = h7(e13)) | (~ spl80_34 | spl80_934)), inference(forward_demodulation, [], [f7028, f1442])).
fof(f1442, plain, ((e13 = op1(e13, e13)) | ~ spl80_34), inference(avatar_component_clause, [], [f1440])).
fof(f1440, plain, (spl80_34 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_34])])).
fof(f7028, plain, (~ (op2(e23, e23) = h7(op1(e13, e13))) | spl80_934), inference(avatar_component_clause, [], [f7026])).
fof(f7026, plain, (spl80_934 <=> (op2(e23, e23) = h7(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_934])])).
fof(f12426, plain, (~ spl80_24 | ~ spl80_149 | ~ spl80_928 | spl80_932), inference(avatar_contradiction_clause, [], [f12425])).
fof(f12425, plain, ($false | (~ spl80_24 | ~ spl80_149 | ~ spl80_928 | spl80_932)), inference(subsumption_resolution, [], [f12424, f1975])).
fof(f1975, plain, ((e23 = op2(e24, e20)) | ~ spl80_149), inference(avatar_component_clause, [], [f1973])).
fof(f1973, plain, (spl80_149 <=> (e23 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_149])])).
fof(f12424, plain, (~ (e23 = op2(e24, e20)) | (~ spl80_24 | ~ spl80_928 | spl80_932)), inference(forward_demodulation, [], [f12423, f732])).
fof(f12423, plain, (~ (op2(e24, e20) = h7(e13)) | (~ spl80_24 | ~ spl80_928 | spl80_932)), inference(forward_demodulation, [], [f12422, f1400])).
fof(f1400, plain, ((e13 = op1(e14, e10)) | ~ spl80_24), inference(avatar_component_clause, [], [f1398])).
fof(f1398, plain, (spl80_24 <=> (e13 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_24])])).
fof(f12422, plain, (~ (op2(e24, e20) = h7(op1(e14, e10))) | (~ spl80_928 | spl80_932)), inference(forward_demodulation, [], [f7020, f7003])).
fof(f7020, plain, (~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | spl80_932), inference(avatar_component_clause, [], [f7018])).
fof(f7018, plain, (spl80_932 <=> (h7(op1(e14, e10)) = op2(h7(e14), e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_932])])).
fof(f12239, plain, (~ spl80_18 | ~ spl80_143 | ~ spl80_464 | ~ spl80_928 | spl80_931), inference(avatar_contradiction_clause, [], [f12238])).
fof(f12238, plain, ($false | (~ spl80_18 | ~ spl80_143 | ~ spl80_464 | ~ spl80_928 | spl80_931)), inference(subsumption_resolution, [], [f12237, f1950])).
fof(f1950, plain, ((e22 = op2(e24, e21)) | ~ spl80_143), inference(avatar_component_clause, [], [f1948])).
fof(f1948, plain, (spl80_143 <=> (e22 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_143])])).
fof(f12237, plain, (~ (e22 = op2(e24, e21)) | (~ spl80_18 | ~ spl80_464 | ~ spl80_928 | spl80_931)), inference(forward_demodulation, [], [f12236, f9385])).
fof(f12236, plain, (~ (op2(e24, e21) = h7(e12)) | (~ spl80_18 | ~ spl80_928 | spl80_931)), inference(forward_demodulation, [], [f12235, f1375])).
fof(f1375, plain, ((e12 = op1(e14, e11)) | ~ spl80_18), inference(avatar_component_clause, [], [f1373])).
fof(f1373, plain, (spl80_18 <=> (e12 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_18])])).
fof(f12235, plain, (~ (op2(e24, e21) = h7(op1(e14, e11))) | (~ spl80_928 | spl80_931)), inference(forward_demodulation, [], [f7016, f7003])).
fof(f7016, plain, (~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | spl80_931), inference(avatar_component_clause, [], [f7014])).
fof(f7014, plain, (spl80_931 <=> (h7(op1(e14, e11)) = op2(h7(e14), e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_931])])).
fof(f12060, plain, (~ spl80_12 | ~ spl80_137 | ~ spl80_464 | ~ spl80_928 | spl80_930), inference(avatar_contradiction_clause, [], [f12059])).
fof(f12059, plain, ($false | (~ spl80_12 | ~ spl80_137 | ~ spl80_464 | ~ spl80_928 | spl80_930)), inference(subsumption_resolution, [], [f12058, f1925])).
fof(f1925, plain, ((e21 = op2(e24, e22)) | ~ spl80_137), inference(avatar_component_clause, [], [f1923])).
fof(f1923, plain, (spl80_137 <=> (e21 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_137])])).
fof(f12058, plain, (~ (e21 = op2(e24, e22)) | (~ spl80_12 | ~ spl80_464 | ~ spl80_928 | spl80_930)), inference(forward_demodulation, [], [f12057, f731])).
fof(f12057, plain, (~ (op2(e24, e22) = h7(e11)) | (~ spl80_12 | ~ spl80_464 | ~ spl80_928 | spl80_930)), inference(forward_demodulation, [], [f12056, f1350])).
fof(f1350, plain, ((e11 = op1(e14, e12)) | ~ spl80_12), inference(avatar_component_clause, [], [f1348])).
fof(f1348, plain, (spl80_12 <=> (e11 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_12])])).
fof(f12056, plain, (~ (op2(e24, e22) = h7(op1(e14, e12))) | (~ spl80_464 | ~ spl80_928 | spl80_930)), inference(forward_demodulation, [], [f12055, f7003])).
fof(f12055, plain, (~ (h7(op1(e14, e12)) = op2(h7(e14), e22)) | (~ spl80_464 | spl80_930)), inference(forward_demodulation, [], [f7012, f3823])).
fof(f7012, plain, (~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | spl80_930), inference(avatar_component_clause, [], [f7010])).
fof(f7010, plain, (spl80_930 <=> (h7(op1(e14, e12)) = op2(h7(e14), h3(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_930])])).
fof(f11893, plain, (~ spl80_6 | spl80_929), inference(avatar_contradiction_clause, [], [f11892])).
fof(f11892, plain, ($false | (~ spl80_6 | spl80_929)), inference(subsumption_resolution, [], [f11891, f2546])).
fof(f11891, plain, (~ (e20 = h7(e10)) | (~ spl80_6 | spl80_929)), inference(forward_demodulation, [], [f7008, f1325])).
fof(f1325, plain, ((e10 = op1(e14, e13)) | ~ spl80_6), inference(avatar_component_clause, [], [f1323])).
fof(f1323, plain, (spl80_6 <=> (e10 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_6])])).
fof(f7008, plain, (~ (e20 = h7(op1(e14, e13))) | spl80_929), inference(avatar_component_clause, [], [f7006])).
fof(f7006, plain, (spl80_929 <=> (e20 = h7(op1(e14, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl80_929])])).
fof(f11773, plain, (~ spl80_115 | ~ spl80_240 | ~ spl80_464 | ~ spl80_928 | spl80_950), inference(avatar_contradiction_clause, [], [f11772])).
fof(f11772, plain, ($false | (~ spl80_115 | ~ spl80_240 | ~ spl80_464 | ~ spl80_928 | spl80_950)), inference(subsumption_resolution, [], [f11771, f2357])).
fof(f2357, plain, ((e24 = op2(e20, e22)) | ~ spl80_240), inference(avatar_component_clause, [], [f2355])).
fof(f2355, plain, (spl80_240 <=> (e24 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_240])])).
fof(f11771, plain, (~ (e24 = op2(e20, e22)) | (~ spl80_115 | ~ spl80_464 | ~ spl80_928 | spl80_950)), inference(forward_demodulation, [], [f11770, f7003])).
fof(f11770, plain, (~ (op2(e20, e22) = h7(e14)) | (~ spl80_115 | ~ spl80_464 | spl80_950)), inference(forward_demodulation, [], [f11769, f1782])).
fof(f1782, plain, ((e14 = op1(e10, e12)) | ~ spl80_115), inference(avatar_component_clause, [], [f1780])).
fof(f1780, plain, (spl80_115 <=> (e14 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_115])])).
fof(f11769, plain, (~ (op2(e20, e22) = h7(op1(e10, e12))) | (~ spl80_464 | spl80_950)), inference(forward_demodulation, [], [f7092, f3823])).
fof(f7092, plain, (~ (h7(op1(e10, e12)) = op2(e20, h3(e14))) | spl80_950), inference(avatar_component_clause, [], [f7090])).
fof(f7090, plain, (spl80_950 <=> (h7(op1(e10, e12)) = op2(e20, h3(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_950])])).
fof(f11734, plain, (~ spl80_125 | ~ spl80_100), inference(avatar_split_clause, [], [f11733, f1717, f1822])).
fof(f1822, plain, (spl80_125 <=> (op1(e10, e10) = e14)), introduced(avatar_definition, [new_symbols(naming, [spl80_125])])).
fof(f11733, plain, (~ (op1(e10, e10) = e14) | ~ spl80_100), inference(forward_demodulation, [], [f400, f1719])).
fof(f400, plain, ~ (op1(e10, e10) = op1(e11, e10)), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (op1(e14, e13) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e14)) & ~ (op1(e14, e12) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e14)) & ~ (op1(e14, e11) = op1(e14, e13)) & ~ (op1(e14, e11) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e14)) & ~ (op1(e14, e10) = op1(e14, e13)) & ~ (op1(e14, e10) = op1(e14, e12)) & ~ (op1(e14, e10) = op1(e14, e11)) & ~ (op1(e13, e13) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e14)) & ~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e14)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e14)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e13) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e14)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e14)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e14)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e13) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e14)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e14)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e14)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e13) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e14)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e14)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e14)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e13, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e14, e14)) & ~ (op1(e12, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e14, e14)) & ~ (op1(e11, e14) = op1(e13, e14)) & ~ (op1(e11, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e14, e14)) & ~ (op1(e10, e14) = op1(e13, e14)) & ~ (op1(e10, e14) = op1(e12, e14)) & ~ (op1(e10, e14) = op1(e11, e14)) & ~ (op1(e13, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e14, e13)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e14, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e14, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e13, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e14, e12)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e14, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e14, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e13, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e14, e11)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e14, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e14, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e13, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e14, e10)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e14, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e14, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax7)).
fof(f11730, plain, (~ spl80_123 | ~ spl80_48), inference(avatar_split_clause, [], [f11729, f1499, f1814])).
fof(f1814, plain, (spl80_123 <=> (op1(e10, e10) = e12)), introduced(avatar_definition, [new_symbols(naming, [spl80_123])])).
fof(f11729, plain, (~ (op1(e10, e10) = e12) | ~ spl80_48), inference(forward_demodulation, [], [f402, f1501])).
fof(f402, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f7])).
fof(f11728, plain, (~ spl80_124 | ~ spl80_24), inference(avatar_split_clause, [], [f11727, f1398, f1818])).
fof(f1818, plain, (spl80_124 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl80_124])])).
fof(f11727, plain, (~ (op1(e10, e10) = e13) | ~ spl80_24), inference(forward_demodulation, [], [f403, f1400])).
fof(f403, plain, ~ (op1(e10, e10) = op1(e14, e10)), inference(cnf_transformation, [], [f7])).
fof(f11722, plain, (~ spl80_122 | ~ spl80_107), inference(avatar_split_clause, [], [f11721, f1747, f1810])).
fof(f1810, plain, (spl80_122 <=> (op1(e10, e10) = e11)), introduced(avatar_definition, [new_symbols(naming, [spl80_122])])).
fof(f11721, plain, (~ (op1(e10, e10) = e11) | ~ spl80_107), inference(forward_demodulation, [], [f452, f1749])).
fof(f452, plain, ~ (op1(e10, e10) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f11668, plain, (~ spl80_201 | ~ spl80_203), inference(avatar_contradiction_clause, [], [f11667])).
fof(f11667, plain, ($false | (~ spl80_201 | ~ spl80_203)), inference(subsumption_resolution, [], [f11666, f611])).
fof(f611, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f10])).
fof(f10, plain, (~ (e23 = e24) & ~ (e22 = e24) & ~ (e22 = e23) & ~ (e21 = e24) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e24) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax10)).
fof(f11666, plain, ((e20 = e22) | (~ spl80_201 | ~ spl80_203)), inference(backward_demodulation, [], [f2202, f2194])).
fof(f2202, plain, ((e22 = op2(e21, e24)) | ~ spl80_203), inference(avatar_component_clause, [], [f2200])).
fof(f2200, plain, (spl80_203 <=> (e22 = op2(e21, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_203])])).
fof(f11639, plain, (~ spl80_216 | ~ spl80_296), inference(avatar_split_clause, [], [f11632, f2919, f2255])).
fof(f2255, plain, (spl80_216 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_216])])).
fof(f2919, plain, (spl80_296 <=> (e20 = h18(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_296])])).
fof(f11632, plain, (~ (e20 = op2(e21, e21)) | ~ spl80_296), inference(backward_demodulation, [], [f2656, f2920])).
fof(f2920, plain, ((e20 = h18(e14)) | ~ spl80_296), inference(avatar_component_clause, [], [f2919])).
fof(f2656, plain, ~ (op2(e21, e21) = h18(e14)), inference(backward_demodulation, [], [f566, f790])).
fof(f790, plain, (op2(e21, e24) = h18(e14)), inference(cnf_transformation, [], [f33])).
fof(f33, plain, ((op2(e21, e24) = h18(e14)) & (h18(e12) = op2(e21, op2(op2(e21, e24), e21))) & (h18(e10) = op2(op2(e21, e24), e21)) & (e21 = h18(e13)) & (e24 = h18(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax33)).
fof(f566, plain, ~ (op2(e21, e21) = op2(e21, e24)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (op2(e24, e23) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e24)) & ~ (op2(e24, e22) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e24)) & ~ (op2(e24, e21) = op2(e24, e23)) & ~ (op2(e24, e21) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e24)) & ~ (op2(e24, e20) = op2(e24, e23)) & ~ (op2(e24, e20) = op2(e24, e22)) & ~ (op2(e24, e20) = op2(e24, e21)) & ~ (op2(e23, e23) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e24)) & ~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e24)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e24)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e23) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e24)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e24)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e24)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e23) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e24)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e24)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e24)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e23) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e24)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e24)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e24)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e23, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e24, e24)) & ~ (op2(e22, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e24, e24)) & ~ (op2(e21, e24) = op2(e23, e24)) & ~ (op2(e21, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e24, e24)) & ~ (op2(e20, e24) = op2(e23, e24)) & ~ (op2(e20, e24) = op2(e22, e24)) & ~ (op2(e20, e24) = op2(e21, e24)) & ~ (op2(e23, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e24, e23)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e24, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e24, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e23, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e24, e22)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e24, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e24, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e23, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e24, e21)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e24, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e24, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e23, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e24, e20)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e24, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e24, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax8)).
fof(f11629, plain, (spl80_185 | ~ spl80_322 | ~ spl80_328 | ~ spl80_504), inference(avatar_split_clause, [], [f11623, f4043, f3089, f3059, f2124])).
fof(f2124, plain, (spl80_185 <=> (e24 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_185])])).
fof(f3059, plain, (spl80_322 <=> (e22 = h16(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_322])])).
fof(f3089, plain, (spl80_328 <=> (e20 = h16(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_328])])).
fof(f4043, plain, (spl80_504 <=> (e23 = h4(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_504])])).
fof(f11623, plain, ((e24 = op2(e22, e23)) | (~ spl80_322 | ~ spl80_328 | ~ spl80_504)), inference(backward_demodulation, [], [f11536, f3060])).
fof(f3060, plain, ((e22 = h16(e10)) | ~ spl80_322), inference(avatar_component_clause, [], [f3059])).
fof(f11536, plain, ((e24 = op2(h16(e10), e23)) | (~ spl80_328 | ~ spl80_504)), inference(backward_demodulation, [], [f10686, f4044])).
fof(f4044, plain, ((e23 = h4(e14)) | ~ spl80_504), inference(avatar_component_clause, [], [f4043])).
fof(f10686, plain, ((e24 = op2(h16(e10), h4(e14))) | ~ spl80_328), inference(backward_demodulation, [], [f2649, f10678])).
fof(f10678, plain, ((h16(e10) = h17(e14)) | ~ spl80_328), inference(backward_demodulation, [], [f785, f10677])).
fof(f10677, plain, ((op2(e20, e24) = h16(e10)) | ~ spl80_328), inference(forward_demodulation, [], [f2640, f3090])).
fof(f3090, plain, ((e20 = h16(e14)) | ~ spl80_328), inference(avatar_component_clause, [], [f3089])).
fof(f2640, plain, (h16(e10) = op2(h16(e14), e24)), inference(forward_demodulation, [], [f778, f780])).
fof(f780, plain, (op2(e24, e23) = h16(e14)), inference(cnf_transformation, [], [f31])).
fof(f31, plain, ((op2(e24, e23) = h16(e14)) & (h16(e12) = op2(e24, op2(op2(e24, e23), e24))) & (h16(e10) = op2(op2(e24, e23), e24)) & (e24 = h16(e13)) & (e23 = h16(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax31)).
fof(f778, plain, (h16(e10) = op2(op2(e24, e23), e24)), inference(cnf_transformation, [], [f31])).
fof(f785, plain, (op2(e20, e24) = h17(e14)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ((op2(e20, e24) = h17(e14)) & (h17(e12) = op2(e20, op2(op2(e20, e24), e20))) & (h17(e10) = op2(op2(e20, e24), e20)) & (e20 = h17(e13)) & (e24 = h17(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax32)).
fof(f2649, plain, (e24 = op2(h17(e14), h4(e14))), inference(backward_demodulation, [], [f2497, f785])).
fof(f2497, plain, (e24 = op2(op2(e20, e24), h4(e14))), inference(backward_demodulation, [], [f395, f720])).
fof(f720, plain, (op2(e24, e20) = h4(e14)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ((op2(e24, e20) = h4(e14)) & (h4(e12) = op2(e24, op2(op2(e24, e20), e24))) & (h4(e10) = op2(op2(e24, e20), e24)) & (e24 = h4(e13)) & (e20 = h4(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax19)).
fof(f395, plain, (e24 = op2(op2(e20, e24), op2(e24, e20))), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ((e24 = op2(op2(e24, e24), op2(e24, e24))) & (e24 = op2(op2(e23, e24), op2(e24, e23))) & (e24 = op2(op2(e22, e24), op2(e24, e22))) & (e24 = op2(op2(e21, e24), op2(e24, e21))) & (e24 = op2(op2(e20, e24), op2(e24, e20))) & (e23 = op2(op2(e24, e23), op2(e23, e24))) & (e23 = op2(op2(e23, e23), op2(e23, e23))) & (e23 = op2(op2(e22, e23), op2(e23, e22))) & (e23 = op2(op2(e21, e23), op2(e23, e21))) & (e23 = op2(op2(e20, e23), op2(e23, e20))) & (e22 = op2(op2(e24, e22), op2(e22, e24))) & (e22 = op2(op2(e23, e22), op2(e22, e23))) & (e22 = op2(op2(e22, e22), op2(e22, e22))) & (e22 = op2(op2(e21, e22), op2(e22, e21))) & (e22 = op2(op2(e20, e22), op2(e22, e20))) & (e21 = op2(op2(e24, e21), op2(e21, e24))) & (e21 = op2(op2(e23, e21), op2(e21, e23))) & (e21 = op2(op2(e22, e21), op2(e21, e22))) & (e21 = op2(op2(e21, e21), op2(e21, e21))) & (e21 = op2(op2(e20, e21), op2(e21, e20))) & (e20 = op2(op2(e24, e20), op2(e20, e24))) & (e20 = op2(op2(e23, e20), op2(e20, e23))) & (e20 = op2(op2(e22, e20), op2(e20, e22))) & (e20 = op2(op2(e21, e20), op2(e20, e21))) & (e20 = op2(op2(e20, e20), op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax6)).
fof(f11627, plain, (~ spl80_178 | ~ spl80_322 | ~ spl80_328), inference(avatar_split_clause, [], [f11617, f3089, f3059, f2095])).
fof(f2095, plain, (spl80_178 <=> (e22 = op2(e22, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_178])])).
fof(f11617, plain, (~ (e22 = op2(e22, e24)) | (~ spl80_322 | ~ spl80_328)), inference(backward_demodulation, [], [f10680, f3060])).
fof(f10680, plain, (~ (op2(e22, e24) = h16(e10)) | ~ spl80_328), inference(backward_demodulation, [], [f2643, f10678])).
fof(f2643, plain, ~ (op2(e22, e24) = h17(e14)), inference(backward_demodulation, [], [f541, f785])).
fof(f541, plain, ~ (op2(e20, e24) = op2(e22, e24)), inference(cnf_transformation, [], [f8])).
fof(f11601, plain, (spl80_143 | ~ spl80_388), inference(avatar_split_clause, [], [f11592, f3411, f1948])).
fof(f3411, plain, (spl80_388 <=> (e21 = h12(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_388])])).
fof(f11592, plain, ((e22 = op2(e24, e21)) | ~ spl80_388), inference(backward_demodulation, [], [f2597, f3412])).
fof(f3412, plain, ((e21 = h12(e14)) | ~ spl80_388), inference(avatar_component_clause, [], [f3411])).
fof(f2597, plain, (e22 = op2(e24, h12(e14))), inference(backward_demodulation, [], [f692, f760])).
fof(f760, plain, (op2(e24, e22) = h12(e14)), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ((op2(e24, e22) = h12(e14)) & (h12(e12) = op2(e24, op2(op2(e24, e22), e24))) & (h12(e10) = op2(op2(e24, e22), e24)) & (e24 = h12(e13)) & (e22 = h12(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax27)).
fof(f692, plain, (e22 = op2(e24, op2(e24, e22))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e24 = op2(e24, op2(e24, e24))) & (e23 = op2(e24, op2(e24, e23))) & (e22 = op2(e24, op2(e24, e22))) & (e21 = op2(e24, op2(e24, e21))) & (e20 = op2(e24, op2(e24, e20))) & (e24 = op2(e23, op2(e23, e24))) & (e23 = op2(e23, op2(e23, e23))) & (e22 = op2(e23, op2(e23, e22))) & (e21 = op2(e23, op2(e23, e21))) & (e20 = op2(e23, op2(e23, e20))) & (e24 = op2(e22, op2(e22, e24))) & (e23 = op2(e22, op2(e22, e23))) & (e22 = op2(e22, op2(e22, e22))) & (e21 = op2(e22, op2(e22, e21))) & (e20 = op2(e22, op2(e22, e20))) & (e24 = op2(e21, op2(e21, e24))) & (e23 = op2(e21, op2(e21, e23))) & (e22 = op2(e21, op2(e21, e22))) & (e21 = op2(e21, op2(e21, e21))) & (e20 = op2(e21, op2(e21, e20))) & (e24 = op2(e20, op2(e20, e24))) & (e23 = op2(e20, op2(e20, e23))) & (e22 = op2(e20, op2(e20, e22))) & (e21 = op2(e20, op2(e20, e21))) & (e20 = op2(e20, op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax13)).
fof(f11572, plain, (spl80_197 | ~ spl80_484), inference(avatar_split_clause, [], [f11559, f3936, f2175])).
fof(f3936, plain, (spl80_484 <=> (e20 = h6(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_484])])).
fof(f11559, plain, ((e21 = op2(e22, e20)) | ~ spl80_484), inference(backward_demodulation, [], [f2529, f3937])).
fof(f3937, plain, ((e20 = h6(e14)) | ~ spl80_484), inference(avatar_component_clause, [], [f3936])).
fof(f2529, plain, (e21 = op2(e22, h6(e14))), inference(backward_demodulation, [], [f681, f730])).
fof(f730, plain, (op2(e22, e21) = h6(e14)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ((op2(e22, e21) = h6(e14)) & (h6(e12) = op2(e22, op2(op2(e22, e21), e22))) & (h6(e10) = op2(op2(e22, e21), e22)) & (e22 = h6(e13)) & (e21 = h6(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax21)).
fof(f681, plain, (e21 = op2(e22, op2(e22, e21))), inference(cnf_transformation, [], [f13])).
fof(f11476, plain, (spl80_228 | ~ spl80_880), inference(avatar_split_clause, [], [f11469, f6615, f2305])).
fof(f6615, plain, (spl80_880 <=> (e24 = h9(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_880])])).
fof(f11469, plain, ((e22 = op2(e20, e24)) | ~ spl80_880), inference(backward_demodulation, [], [f2564, f6616])).
fof(f6616, plain, ((e24 = h9(e14)) | ~ spl80_880), inference(avatar_component_clause, [], [f6615])).
fof(f2564, plain, (e22 = op2(e20, h9(e14))), inference(backward_demodulation, [], [f672, f745])).
fof(f745, plain, (op2(e20, e22) = h9(e14)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ((op2(e20, e22) = h9(e14)) & (h9(e12) = op2(e20, op2(op2(e20, e22), e20))) & (h9(e10) = op2(op2(e20, e22), e20)) & (e20 = h9(e13)) & (e22 = h9(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax24)).
fof(f672, plain, (e22 = op2(e20, op2(e20, e22))), inference(cnf_transformation, [], [f13])).
fof(f11458, plain, (~ spl80_27 | spl80_933), inference(avatar_contradiction_clause, [], [f11457])).
fof(f11457, plain, ($false | (~ spl80_27 | spl80_933)), inference(subsumption_resolution, [], [f11456, f731])).
fof(f11456, plain, (~ (e21 = h7(e11)) | (~ spl80_27 | spl80_933)), inference(forward_demodulation, [], [f7024, f1413])).
fof(f1413, plain, ((e11 = op1(e13, e14)) | ~ spl80_27), inference(avatar_component_clause, [], [f1411])).
fof(f1411, plain, (spl80_27 <=> (e11 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_27])])).
fof(f7024, plain, (~ (e21 = h7(op1(e13, e14))) | spl80_933), inference(avatar_component_clause, [], [f7022])).
fof(f7022, plain, (spl80_933 <=> (e21 = h7(op1(e13, e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_933])])).
fof(f11449, plain, (spl80_179 | ~ spl80_464 | ~ spl80_1064), inference(avatar_split_clause, [], [f11438, f7917, f3822, f2099])).
fof(f11438, plain, ((e23 = op2(e22, e24)) | (~ spl80_464 | ~ spl80_1064)), inference(backward_demodulation, [], [f9398, f7918])).
fof(f9398, plain, ((e23 = op2(e22, h3(e10))) | ~ spl80_464), inference(backward_demodulation, [], [f2628, f9392])).
fof(f9392, plain, ((h3(e10) = h15(e14)) | ~ spl80_464), inference(backward_demodulation, [], [f775, f9384])).
fof(f9384, plain, ((op2(e22, e23) = h3(e10)) | ~ spl80_464), inference(backward_demodulation, [], [f2494, f3823])).
fof(f2494, plain, (h3(e10) = op2(h3(e14), e23)), inference(forward_demodulation, [], [f713, f715])).
fof(f713, plain, (h3(e10) = op2(op2(e23, e20), e23)), inference(cnf_transformation, [], [f18])).
fof(f775, plain, (op2(e22, e23) = h15(e14)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ((op2(e22, e23) = h15(e14)) & (h15(e12) = op2(e22, op2(op2(e22, e23), e22))) & (h15(e10) = op2(op2(e22, e23), e22)) & (e22 = h15(e13)) & (e23 = h15(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax30)).
fof(f2628, plain, (e23 = op2(e22, h15(e14))), inference(backward_demodulation, [], [f683, f775])).
fof(f683, plain, (e23 = op2(e22, op2(e22, e23))), inference(cnf_transformation, [], [f13])).
fof(f11400, plain, (spl80_1102 | ~ spl80_225), inference(avatar_split_clause, [], [f11399, f2292, f8167])).
fof(f2292, plain, (spl80_225 <=> (e24 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_225])])).
fof(f11399, plain, ((e24 = h1(e14)) | ~ spl80_225), inference(forward_demodulation, [], [f705, f2294])).
fof(f2294, plain, ((e24 = op2(e21, e20)) | ~ spl80_225), inference(avatar_component_clause, [], [f2292])).
fof(f705, plain, (op2(e21, e20) = h1(e14)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(e21, e20) = h1(e14)) & (h1(e12) = op2(e21, op2(op2(e21, e20), e21))) & (h1(e10) = op2(op2(e21, e20), e21)) & (e21 = h1(e13)) & (e20 = h1(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax16)).
fof(f11344, plain, (spl80_158 | ~ spl80_189), inference(avatar_contradiction_clause, [], [f11343])).
fof(f11343, plain, ($false | (spl80_158 | ~ spl80_189)), inference(subsumption_resolution, [], [f11338, f2012])).
fof(f2012, plain, (~ (e22 = op2(e23, e23)) | spl80_158), inference(avatar_component_clause, [], [f2011])).
fof(f2011, plain, (spl80_158 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_158])])).
fof(f11338, plain, ((e22 = op2(e23, e23)) | ~ spl80_189), inference(backward_demodulation, [], [f387, f2143])).
fof(f2143, plain, ((e23 = op2(e22, e22)) | ~ spl80_189), inference(avatar_component_clause, [], [f2141])).
fof(f2141, plain, (spl80_189 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_189])])).
fof(f387, plain, (e22 = op2(op2(e22, e22), op2(e22, e22))), inference(cnf_transformation, [], [f6])).
fof(f11278, plain, (spl80_156 | ~ spl80_249), inference(avatar_contradiction_clause, [], [f11277])).
fof(f11277, plain, ($false | (spl80_156 | ~ spl80_249)), inference(subsumption_resolution, [], [f11270, f2004])).
fof(f2004, plain, (~ (e20 = op2(e23, e23)) | spl80_156), inference(avatar_component_clause, [], [f2003])).
fof(f2003, plain, (spl80_156 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_156])])).
fof(f11270, plain, ((e20 = op2(e23, e23)) | ~ spl80_249), inference(backward_demodulation, [], [f375, f2395])).
fof(f2395, plain, ((op2(e20, e20) = e23) | ~ spl80_249), inference(avatar_component_clause, [], [f2393])).
fof(f2393, plain, (spl80_249 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl80_249])])).
fof(f375, plain, (e20 = op2(op2(e20, e20), op2(e20, e20))), inference(cnf_transformation, [], [f6])).
fof(f11201, plain, (spl80_232 | ~ spl80_488), inference(avatar_split_clause, [], [f11193, f3957, f2322])).
fof(f2322, plain, (spl80_232 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_232])])).
fof(f11193, plain, ((e21 = op2(e20, e23)) | ~ spl80_488), inference(backward_demodulation, [], [f2515, f3958])).
fof(f2515, plain, (e21 = op2(e20, h5(e14))), inference(backward_demodulation, [], [f671, f725])).
fof(f725, plain, (op2(e20, e21) = h5(e14)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ((op2(e20, e21) = h5(e14)) & (h5(e12) = op2(e20, op2(op2(e20, e21), e20))) & (h5(e10) = op2(op2(e20, e21), e20)) & (e20 = h5(e13)) & (e21 = h5(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax20)).
fof(f671, plain, (e21 = op2(e20, op2(e20, e21))), inference(cnf_transformation, [], [f13])).
fof(f11079, plain, (spl80_128 | ~ spl80_796), inference(avatar_contradiction_clause, [], [f11078])).
fof(f11078, plain, ($false | (spl80_128 | ~ spl80_796)), inference(subsumption_resolution, [], [f11073, f1886])).
fof(f1886, plain, (~ (e22 = op2(e24, e24)) | spl80_128), inference(avatar_component_clause, [], [f1885])).
fof(f1885, plain, (spl80_128 <=> (e22 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_128])])).
fof(f11073, plain, ((e22 = op2(e24, e24)) | ~ spl80_796), inference(backward_demodulation, [], [f2597, f6034])).
fof(f6034, plain, ((e24 = h12(e14)) | ~ spl80_796), inference(avatar_component_clause, [], [f6033])).
fof(f6033, plain, (spl80_796 <=> (e24 = h12(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_796])])).
fof(f11068, plain, (spl80_201 | ~ spl80_1102), inference(avatar_split_clause, [], [f11062, f8167, f2192])).
fof(f11062, plain, ((e20 = op2(e21, e24)) | ~ spl80_1102), inference(backward_demodulation, [], [f2467, f8168])).
fof(f2467, plain, (e20 = op2(e21, h1(e14))), inference(backward_demodulation, [], [f675, f705])).
fof(f675, plain, (e20 = op2(e21, op2(e21, e20))), inference(cnf_transformation, [], [f13])).
fof(f11024, plain, (spl80_218 | ~ spl80_420), inference(avatar_split_clause, [], [f11023, f3582, f2263])).
fof(f2263, plain, (spl80_218 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_218])])).
fof(f3582, plain, (spl80_420 <=> (e21 = h10(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_420])])).
fof(f11023, plain, ((e22 = op2(e21, e21)) | ~ spl80_420), inference(forward_demodulation, [], [f2576, f3583])).
fof(f3583, plain, ((e21 = h10(e14)) | ~ spl80_420), inference(avatar_component_clause, [], [f3582])).
fof(f2576, plain, (e22 = op2(e21, h10(e14))), inference(backward_demodulation, [], [f677, f750])).
fof(f750, plain, (op2(e21, e22) = h10(e14)), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ((op2(e21, e22) = h10(e14)) & (h10(e12) = op2(e21, op2(op2(e21, e22), e21))) & (h10(e10) = op2(op2(e21, e22), e21)) & (e21 = h10(e13)) & (e22 = h10(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax25)).
fof(f677, plain, (e22 = op2(e21, op2(e21, e22))), inference(cnf_transformation, [], [f13])).
fof(f11001, plain, (spl80_149 | ~ spl80_328), inference(avatar_split_clause, [], [f10688, f3089, f1973])).
fof(f10688, plain, ((e23 = op2(e24, e20)) | ~ spl80_328), inference(forward_demodulation, [], [f2638, f3090])).
fof(f2638, plain, (e23 = op2(e24, h16(e14))), inference(backward_demodulation, [], [f693, f780])).
fof(f693, plain, (e23 = op2(e24, op2(e24, e23))), inference(cnf_transformation, [], [f13])).
fof(f10953, plain, (~ spl80_83 | ~ spl80_93), inference(avatar_split_clause, [], [f10945, f1688, f1646])).
fof(f1688, plain, (spl80_93 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_93])])).
fof(f10945, plain, (~ (e12 = op1(e11, e13)) | ~ spl80_93), inference(backward_demodulation, [], [f465, f1690])).
fof(f1690, plain, ((e12 = op1(e11, e11)) | ~ spl80_93), inference(avatar_component_clause, [], [f1688])).
fof(f465, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f10881, plain, (~ spl80_198 | spl80_536), inference(avatar_contradiction_clause, [], [f10880])).
fof(f10880, plain, ($false | (~ spl80_198 | spl80_536)), inference(subsumption_resolution, [], [f10878, f4221])).
fof(f4221, plain, (~ (e22 = h2(e14)) | spl80_536), inference(avatar_component_clause, [], [f4219])).
fof(f4219, plain, (spl80_536 <=> (e22 = h2(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_536])])).
fof(f10878, plain, ((e22 = h2(e14)) | ~ spl80_198), inference(backward_demodulation, [], [f710, f2181])).
fof(f2181, plain, ((e22 = op2(e22, e20)) | ~ spl80_198), inference(avatar_component_clause, [], [f2179])).
fof(f2179, plain, (spl80_198 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_198])])).
fof(f710, plain, (op2(e22, e20) = h2(e14)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(e22, e20) = h2(e14)) & (h2(e12) = op2(e22, op2(op2(e22, e20), e22))) & (h2(e10) = op2(op2(e22, e20), e22)) & (e22 = h2(e13)) & (e20 = h2(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax17)).
fof(f10858, plain, (spl80_187 | ~ spl80_218), inference(avatar_split_clause, [], [f10849, f2263, f2133])).
fof(f2133, plain, (spl80_187 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_187])])).
fof(f10849, plain, ((e21 = op2(e22, e22)) | ~ spl80_218), inference(backward_demodulation, [], [f381, f2265])).
fof(f2265, plain, ((e22 = op2(e21, e21)) | ~ spl80_218), inference(avatar_component_clause, [], [f2263])).
fof(f381, plain, (e21 = op2(op2(e21, e21), op2(e21, e21))), inference(cnf_transformation, [], [f6])).
fof(f10845, plain, (spl80_322 | ~ spl80_228 | ~ spl80_328), inference(avatar_split_clause, [], [f10843, f3089, f2305, f3059])).
fof(f10843, plain, ((e22 = h16(e10)) | (~ spl80_228 | ~ spl80_328)), inference(backward_demodulation, [], [f10677, f2307])).
fof(f10832, plain, (spl80_488 | ~ spl80_244), inference(avatar_split_clause, [], [f10831, f2372, f3957])).
fof(f2372, plain, (spl80_244 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_244])])).
fof(f10831, plain, ((e23 = h5(e14)) | ~ spl80_244), inference(backward_demodulation, [], [f725, f2374])).
fof(f2374, plain, ((e23 = op2(e20, e21)) | ~ spl80_244), inference(avatar_component_clause, [], [f2372])).
fof(f10820, plain, (spl80_126 | ~ spl80_250), inference(avatar_contradiction_clause, [], [f10819])).
fof(f10819, plain, ($false | (spl80_126 | ~ spl80_250)), inference(subsumption_resolution, [], [f10811, f1878])).
fof(f1878, plain, (~ (e20 = op2(e24, e24)) | spl80_126), inference(avatar_component_clause, [], [f1877])).
fof(f1877, plain, (spl80_126 <=> (e20 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_126])])).
fof(f10811, plain, ((e20 = op2(e24, e24)) | ~ spl80_250), inference(backward_demodulation, [], [f375, f2399])).
fof(f2399, plain, ((op2(e20, e20) = e24) | ~ spl80_250), inference(avatar_component_clause, [], [f2397])).
fof(f2397, plain, (spl80_250 <=> (op2(e20, e20) = e24)), introduced(avatar_definition, [new_symbols(naming, [spl80_250])])).
fof(f10810, plain, (spl80_244 | ~ spl80_260 | ~ spl80_328), inference(avatar_split_clause, [], [f10804, f3089, f2728, f2372])).
fof(f2728, plain, (spl80_260 <=> (e21 = h20(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_260])])).
fof(f10804, plain, ((e23 = op2(e20, e21)) | (~ spl80_260 | ~ spl80_328)), inference(backward_demodulation, [], [f10676, f2729])).
fof(f2729, plain, ((e21 = h20(e14)) | ~ spl80_260), inference(avatar_component_clause, [], [f2728])).
fof(f10676, plain, ((e23 = op2(e20, h20(e14))) | ~ spl80_328), inference(forward_demodulation, [], [f2675, f3090])).
fof(f2675, plain, (e23 = op2(h16(e14), h20(e14))), inference(backward_demodulation, [], [f2634, f800])).
fof(f800, plain, (op2(e23, e24) = h20(e14)), inference(cnf_transformation, [], [f35])).
fof(f35, plain, ((op2(e23, e24) = h20(e14)) & (h20(e12) = op2(e23, op2(op2(e23, e24), e23))) & (h20(e10) = op2(op2(e23, e24), e23)) & (e23 = h20(e13)) & (e24 = h20(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax35)).
fof(f2634, plain, (e23 = op2(h16(e14), op2(e23, e24))), inference(backward_demodulation, [], [f394, f780])).
fof(f394, plain, (e23 = op2(op2(e24, e23), op2(e23, e24))), inference(cnf_transformation, [], [f6])).
fof(f10809, plain, (spl80_225 | ~ spl80_260 | ~ spl80_328), inference(avatar_split_clause, [], [f10803, f3089, f2728, f2292])).
fof(f10803, plain, ((e24 = op2(e21, e20)) | (~ spl80_260 | ~ spl80_328)), inference(backward_demodulation, [], [f10675, f2729])).
fof(f10675, plain, ((e24 = op2(h20(e14), e20)) | ~ spl80_328), inference(forward_demodulation, [], [f2676, f3090])).
fof(f2676, plain, (e24 = op2(h20(e14), h16(e14))), inference(backward_demodulation, [], [f2635, f800])).
fof(f2635, plain, (e24 = op2(op2(e23, e24), h16(e14))), inference(backward_demodulation, [], [f398, f780])).
fof(f398, plain, (e24 = op2(op2(e23, e24), op2(e24, e23))), inference(cnf_transformation, [], [f6])).
fof(f10791, plain, (spl80_214 | ~ spl80_372 | ~ spl80_464), inference(avatar_split_clause, [], [f10783, f3822, f3325, f2246])).
fof(f10783, plain, ((e23 = op2(e21, e22)) | (~ spl80_372 | ~ spl80_464)), inference(backward_demodulation, [], [f9387, f3326])).
fof(f9387, plain, ((e23 = op2(h13(e14), e22)) | ~ spl80_464), inference(backward_demodulation, [], [f2609, f3823])).
fof(f2609, plain, (e23 = op2(h13(e14), h3(e14))), inference(backward_demodulation, [], [f2485, f765])).
fof(f765, plain, (op2(e20, e23) = h13(e14)), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ((op2(e20, e23) = h13(e14)) & (h13(e12) = op2(e20, op2(op2(e20, e23), e20))) & (h13(e10) = op2(op2(e20, e23), e20)) & (e20 = h13(e13)) & (e23 = h13(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax28)).
fof(f2485, plain, (e23 = op2(op2(e20, e23), h3(e14))), inference(backward_demodulation, [], [f390, f715])).
fof(f390, plain, (e23 = op2(op2(e20, e23), op2(e23, e20))), inference(cnf_transformation, [], [f6])).
fof(f10790, plain, (spl80_191 | ~ spl80_372 | ~ spl80_464), inference(avatar_contradiction_clause, [], [f10789])).
fof(f10789, plain, ($false | (spl80_191 | ~ spl80_372 | ~ spl80_464)), inference(subsumption_resolution, [], [f10782, f2151])).
fof(f2151, plain, (~ (e20 = op2(e22, e21)) | spl80_191), inference(avatar_component_clause, [], [f2150])).
fof(f10782, plain, ((e20 = op2(e22, e21)) | (~ spl80_372 | ~ spl80_464)), inference(backward_demodulation, [], [f9386, f3326])).
fof(f9386, plain, ((e20 = op2(e22, h13(e14))) | ~ spl80_464), inference(backward_demodulation, [], [f2608, f3823])).
fof(f2608, plain, (e20 = op2(h3(e14), h13(e14))), inference(backward_demodulation, [], [f2484, f765])).
fof(f2484, plain, (e20 = op2(h3(e14), op2(e20, e23))), inference(backward_demodulation, [], [f378, f715])).
fof(f378, plain, (e20 = op2(op2(e23, e20), op2(e20, e23))), inference(cnf_transformation, [], [f6])).
fof(f10784, plain, (~ spl80_247 | ~ spl80_372), inference(avatar_split_clause, [], [f10776, f3325, f2385])).
fof(f2385, plain, (spl80_247 <=> (op2(e20, e20) = e21)), introduced(avatar_definition, [new_symbols(naming, [spl80_247])])).
fof(f10776, plain, (~ (op2(e20, e20) = e21) | ~ spl80_372), inference(backward_demodulation, [], [f2605, f3326])).
fof(f2605, plain, ~ (op2(e20, e20) = h13(e14)), inference(backward_demodulation, [], [f552, f765])).
fof(f552, plain, ~ (op2(e20, e20) = op2(e20, e23)), inference(cnf_transformation, [], [f8])).
fof(f10713, plain, (spl80_189 | ~ spl80_464 | ~ spl80_524), inference(avatar_split_clause, [], [f10702, f4154, f3822, f2141])).
fof(f4154, plain, (spl80_524 <=> (e22 = h3(e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_524])])).
fof(f10702, plain, ((e23 = op2(e22, e22)) | (~ spl80_464 | ~ spl80_524)), inference(backward_demodulation, [], [f9398, f4155])).
fof(f4155, plain, ((e22 = h3(e10)) | ~ spl80_524), inference(avatar_component_clause, [], [f4154])).
fof(f10671, plain, (spl80_128 | ~ spl80_190), inference(avatar_contradiction_clause, [], [f10670])).
fof(f10670, plain, ($false | (spl80_128 | ~ spl80_190)), inference(subsumption_resolution, [], [f10669, f1886])).
fof(f10669, plain, ((e22 = op2(e24, e24)) | ~ spl80_190), inference(forward_demodulation, [], [f387, f2147])).
fof(f2147, plain, ((e24 = op2(e22, e22)) | ~ spl80_190), inference(avatar_component_clause, [], [f2145])).
fof(f2145, plain, (spl80_190 <=> (e24 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_190])])).
fof(f10654, plain, (spl80_152 | ~ spl80_928), inference(avatar_split_clause, [], [f9371, f7002, f1986])).
fof(f1986, plain, (spl80_152 <=> (e21 = op2(e23, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_152])])).
fof(f9371, plain, ((e21 = op2(e23, e24)) | ~ spl80_928), inference(backward_demodulation, [], [f2540, f7003])).
fof(f2540, plain, (e21 = op2(e23, h7(e14))), inference(backward_demodulation, [], [f686, f735])).
fof(f686, plain, (e21 = op2(e23, op2(e23, e21))), inference(cnf_transformation, [], [f13])).
fof(f10645, plain, (~ spl80_130 | ~ spl80_5 | spl80_927 | ~ spl80_928), inference(avatar_split_clause, [], [f10644, f7002, f6998, f1318, f1893])).
fof(f1893, plain, (spl80_130 <=> (e24 = op2(e24, e24))), introduced(avatar_definition, [new_symbols(naming, [spl80_130])])).
fof(f1318, plain, (spl80_5 <=> (e14 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_5])])).
fof(f6998, plain, (spl80_927 <=> (h7(op1(e14, e14)) = op2(h7(e14), h7(e14)))), introduced(avatar_definition, [new_symbols(naming, [spl80_927])])).
fof(f10644, plain, (~ (e24 = op2(e24, e24)) | (~ spl80_5 | spl80_927 | ~ spl80_928)), inference(forward_demodulation, [], [f10643, f7003])).
fof(f10643, plain, (~ (op2(e24, e24) = h7(e14)) | (~ spl80_5 | spl80_927 | ~ spl80_928)), inference(forward_demodulation, [], [f9770, f1320])).
fof(f1320, plain, ((e14 = op1(e14, e14)) | ~ spl80_5), inference(avatar_component_clause, [], [f1318])).
fof(f9770, plain, (~ (op2(e24, e24) = h7(op1(e14, e14))) | (spl80_927 | ~ spl80_928)), inference(forward_demodulation, [], [f7000, f7003])).
fof(f7000, plain, (~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | spl80_927), inference(avatar_component_clause, [], [f6998])).
fof(f10635, plain, (~ spl80_84 | ~ spl80_89), inference(avatar_split_clause, [], [f10634, f1671, f1650])).
fof(f1650, plain, (spl80_84 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_84])])).
fof(f10634, plain, (~ (e13 = op1(e11, e13)) | ~ spl80_89), inference(forward_demodulation, [], [f467, f1673])).
fof(f467, plain, ~ (op1(e11, e12) = op1(e11, e13)), inference(cnf_transformation, [], [f7])).
fof(f10616, plain, (spl80_83 | ~ spl80_89), inference(avatar_split_clause, [], [f10615, f1671, f1646])).
fof(f10615, plain, ((e12 = op1(e11, e13)) | ~ spl80_89), inference(forward_demodulation, [], [f652, f1673])).
fof(f652, plain, (e12 = op1(e11, op1(e11, e12))), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e14 = op1(e14, op1(e14, e14))) & (e13 = op1(e14, op1(e14, e13))) & (e12 = op1(e14, op1(e14, e12))) & (e11 = op1(e14, op1(e14, e11))) & (e10 = op1(e14, op1(e14, e10))) & (e14 = op1(e13, op1(e13, e14))) & (e13 = op1(e13, op1(e13, e13))) & (e12 = op1(e13, op1(e13, e12))) & (e11 = op1(e13, op1(e13, e11))) & (e10 = op1(e13, op1(e13, e10))) & (e14 = op1(e12, op1(e12, e14))) & (e13 = op1(e12, op1(e12, e13))) & (e12 = op1(e12, op1(e12, e12))) & (e11 = op1(e12, op1(e12, e11))) & (e10 = op1(e12, op1(e12, e10))) & (e14 = op1(e11, op1(e11, e14))) & (e13 = op1(e11, op1(e11, e13))) & (e12 = op1(e11, op1(e11, e12))) & (e11 = op1(e11, op1(e11, e11))) & (e10 = op1(e11, op1(e11, e10))) & (e14 = op1(e10, op1(e10, e14))) & (e13 = op1(e10, op1(e10, e13))) & (e12 = op1(e10, op1(e10, e12))) & (e11 = op1(e10, op1(e10, e11))) & (e10 = op1(e10, op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax12)).
fof(f10602, plain, (spl80_18 | ~ spl80_12), inference(avatar_split_clause, [], [f10254, f1348, f1373])).
fof(f10254, plain, ((e12 = op1(e14, e11)) | ~ spl80_12), inference(forward_demodulation, [], [f667, f1350])).
fof(f667, plain, (e12 = op1(e14, op1(e14, e12))), inference(cnf_transformation, [], [f12])).
fof(f10550, plain, (spl80_66 | ~ spl80_48 | ~ spl80_107), inference(avatar_split_clause, [], [f10547, f1747, f1499, f1575])).
fof(f10547, plain, ((e10 = op1(e12, e11)) | (~ spl80_48 | ~ spl80_107)), inference(backward_demodulation, [], [f9250, f1749])).
fof(f9250, plain, ((e10 = op1(e12, op1(e10, e13))) | ~ spl80_48), inference(forward_demodulation, [], [f278, f1501])).
fof(f278, plain, (e10 = op1(op1(e13, e10), op1(e10, e13))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, ((e14 = op1(op1(e14, e14), op1(e14, e14))) & (e14 = op1(op1(e13, e14), op1(e14, e13))) & (e14 = op1(op1(e12, e14), op1(e14, e12))) & (e14 = op1(op1(e11, e14), op1(e14, e11))) & (e14 = op1(op1(e10, e14), op1(e14, e10))) & (e13 = op1(op1(e14, e13), op1(e13, e14))) & (e13 = op1(op1(e13, e13), op1(e13, e13))) & (e13 = op1(op1(e12, e13), op1(e13, e12))) & (e13 = op1(op1(e11, e13), op1(e13, e11))) & (e13 = op1(op1(e10, e13), op1(e13, e10))) & (e12 = op1(op1(e14, e12), op1(e12, e14))) & (e12 = op1(op1(e13, e12), op1(e12, e13))) & (e12 = op1(op1(e12, e12), op1(e12, e12))) & (e12 = op1(op1(e11, e12), op1(e12, e11))) & (e12 = op1(op1(e10, e12), op1(e12, e10))) & (e11 = op1(op1(e14, e11), op1(e11, e14))) & (e11 = op1(op1(e13, e11), op1(e11, e13))) & (e11 = op1(op1(e12, e11), op1(e11, e12))) & (e11 = op1(op1(e11, e11), op1(e11, e11))) & (e11 = op1(op1(e10, e11), op1(e11, e10))) & (e10 = op1(op1(e14, e10), op1(e10, e14))) & (e10 = op1(op1(e13, e10), op1(e10, e13))) & (e10 = op1(op1(e12, e10), op1(e10, e12))) & (e10 = op1(op1(e11, e10), op1(e10, e11))) & (e10 = op1(op1(e10, e10), op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax3)).
fof(f10549, plain, (spl80_89 | ~ spl80_48 | ~ spl80_107), inference(avatar_split_clause, [], [f10546, f1747, f1499, f1671])).
fof(f10546, plain, ((e13 = op1(e11, e12)) | (~ spl80_48 | ~ spl80_107)), inference(backward_demodulation, [], [f9249, f1749])).
fof(f9249, plain, ((e13 = op1(op1(e10, e13), e12)) | ~ spl80_48), inference(forward_demodulation, [], [f290, f1501])).
fof(f290, plain, (e13 = op1(op1(e10, e13), op1(e13, e10))), inference(cnf_transformation, [], [f3])).
fof(f10486, plain, (~ spl80_193 | spl80_476), inference(avatar_contradiction_clause, [], [f10485])).
fof(f10485, plain, ($false | (~ spl80_193 | spl80_476)), inference(subsumption_resolution, [], [f10482, f3894])).
fof(f3894, plain, (~ (e22 = h6(e14)) | spl80_476), inference(avatar_component_clause, [], [f3892])).
fof(f3892, plain, (spl80_476 <=> (e22 = h6(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_476])])).
fof(f10482, plain, ((e22 = h6(e14)) | ~ spl80_193), inference(backward_demodulation, [], [f730, f2160])).
fof(f2160, plain, ((e22 = op2(e22, e21)) | ~ spl80_193), inference(avatar_component_clause, [], [f2158])).
fof(f2158, plain, (spl80_193 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_193])])).
fof(f10468, plain, (~ spl80_213 | spl80_416), inference(avatar_contradiction_clause, [], [f10467])).
fof(f10467, plain, ($false | (~ spl80_213 | spl80_416)), inference(subsumption_resolution, [], [f10464, f3562])).
fof(f3562, plain, (~ (e22 = h10(e14)) | spl80_416), inference(avatar_component_clause, [], [f3560])).
fof(f3560, plain, (spl80_416 <=> (e22 = h10(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_416])])).
fof(f10464, plain, ((e22 = h10(e14)) | ~ spl80_213), inference(backward_demodulation, [], [f750, f2244])).
fof(f2244, plain, ((e22 = op2(e21, e22)) | ~ spl80_213), inference(avatar_component_clause, [], [f2242])).
fof(f2242, plain, (spl80_213 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_213])])).
fof(f10451, plain, (spl80_157 | ~ spl80_219), inference(avatar_split_clause, [], [f10442, f2267, f2007])).
fof(f2007, plain, (spl80_157 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_157])])).
fof(f2267, plain, (spl80_219 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_219])])).
fof(f10442, plain, ((e21 = op2(e23, e23)) | ~ spl80_219), inference(backward_demodulation, [], [f381, f2269])).
fof(f2269, plain, ((e23 = op2(e21, e21)) | ~ spl80_219), inference(avatar_component_clause, [], [f2267])).
fof(f10370, plain, (~ spl80_248 | ~ spl80_464), inference(avatar_split_clause, [], [f10369, f3822, f2389])).
fof(f2389, plain, (spl80_248 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl80_248])])).
fof(f10369, plain, (~ (op2(e20, e20) = e22) | ~ spl80_464), inference(forward_demodulation, [], [f2486, f3823])).
fof(f2486, plain, ~ (op2(e20, e20) = h3(e14)), inference(backward_demodulation, [], [f502, f715])).
fof(f502, plain, ~ (op2(e20, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f8])).
fof(f10368, plain, (~ spl80_436 | ~ spl80_232), inference(avatar_split_clause, [], [f10367, f2322, f3668])).
fof(f3668, plain, (spl80_436 <=> (e21 = h9(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_436])])).
fof(f10367, plain, (~ (e21 = h9(e14)) | ~ spl80_232), inference(forward_demodulation, [], [f2562, f2324])).
fof(f2324, plain, ((e21 = op2(e20, e23)) | ~ spl80_232), inference(avatar_component_clause, [], [f2322])).
fof(f2562, plain, ~ (op2(e20, e23) = h9(e14)), inference(backward_demodulation, [], [f557, f745])).
fof(f557, plain, ~ (op2(e20, e22) = op2(e20, e23)), inference(cnf_transformation, [], [f8])).
fof(f10365, plain, (spl80_372 | ~ spl80_232), inference(avatar_split_clause, [], [f10364, f2322, f3325])).
fof(f10364, plain, ((e21 = h13(e14)) | ~ spl80_232), inference(forward_demodulation, [], [f765, f2324])).
fof(f10325, plain, (~ spl80_145 | ~ spl80_928), inference(avatar_split_clause, [], [f9906, f7002, f1956])).
fof(f1956, plain, (spl80_145 <=> (e24 = op2(e24, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_145])])).
fof(f9906, plain, (~ (e24 = op2(e24, e21)) | ~ spl80_928), inference(forward_demodulation, [], [f2536, f7003])).
fof(f2536, plain, ~ (op2(e24, e21) = h7(e14)), inference(backward_demodulation, [], [f519, f735])).
fof(f519, plain, ~ (op2(e23, e21) = op2(e24, e21)), inference(cnf_transformation, [], [f8])).
fof(f10272, plain, (~ spl80_109 | ~ spl80_119), inference(avatar_split_clause, [], [f10271, f1797, f1755])).
fof(f1755, plain, (spl80_109 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_109])])).
fof(f10271, plain, (~ (e13 = op1(e10, e13)) | ~ spl80_119), inference(forward_demodulation, [], [f455, f1799])).
fof(f455, plain, ~ (op1(e10, e11) = op1(e10, e13)), inference(cnf_transformation, [], [f7])).
fof(f10241, plain, (~ spl80_4 | ~ spl80_24), inference(avatar_split_clause, [], [f10240, f1398, f1314])).
fof(f1314, plain, (spl80_4 <=> (e13 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_4])])).
fof(f10240, plain, (~ (e13 = op1(e14, e14)) | ~ spl80_24), inference(forward_demodulation, [], [f493, f1400])).
fof(f493, plain, ~ (op1(e14, e10) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f10239, plain, (~ spl80_3 | ~ spl80_18), inference(avatar_split_clause, [], [f10238, f1373, f1310])).
fof(f1310, plain, (spl80_3 <=> (e12 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_3])])).
fof(f10238, plain, (~ (e12 = op1(e14, e14)) | ~ spl80_18), inference(forward_demodulation, [], [f496, f1375])).
fof(f496, plain, ~ (op1(e14, e11) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f10236, plain, (~ spl80_1 | ~ spl80_6), inference(avatar_split_clause, [], [f10235, f1323, f1302])).
fof(f1302, plain, (spl80_1 <=> (e10 = op1(e14, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_1])])).
fof(f10235, plain, (~ (e10 = op1(e14, e14)) | ~ spl80_6), inference(forward_demodulation, [], [f499, f1325])).
fof(f499, plain, ~ (op1(e14, e13) = op1(e14, e14)), inference(cnf_transformation, [], [f7])).
fof(f10228, plain, (~ spl80_6 | ~ spl80_14), inference(avatar_contradiction_clause, [], [f10227])).
fof(f10227, plain, ($false | (~ spl80_6 | ~ spl80_14)), inference(subsumption_resolution, [], [f10226, f601])).
fof(f601, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f9])).
fof(f9, plain, (~ (e13 = e14) & ~ (e12 = e14) & ~ (e12 = e13) & ~ (e11 = e14) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e14) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax9)).
fof(f10226, plain, ((e10 = e12) | (~ spl80_6 | ~ spl80_14)), inference(forward_demodulation, [], [f10218, f1325])).
fof(f10218, plain, ((e12 = op1(e14, e13)) | ~ spl80_14), inference(backward_demodulation, [], [f667, f1358])).
fof(f1358, plain, ((e13 = op1(e14, e12)) | ~ spl80_14), inference(avatar_component_clause, [], [f1356])).
fof(f1356, plain, (spl80_14 <=> (e13 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_14])])).
fof(f10190, plain, (~ spl80_36 | ~ spl80_48 | ~ spl80_59), inference(avatar_contradiction_clause, [], [f10189])).
fof(f10189, plain, ($false | (~ spl80_36 | ~ spl80_48 | ~ spl80_59)), inference(subsumption_resolution, [], [f10188, f607])).
fof(f607, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f9])).
fof(f10188, plain, ((e12 = e13) | (~ spl80_36 | ~ spl80_48 | ~ spl80_59)), inference(forward_demodulation, [], [f10184, f1501])).
fof(f10184, plain, ((e13 = op1(e13, e10)) | (~ spl80_36 | ~ spl80_59)), inference(backward_demodulation, [], [f9154, f1547])).
fof(f1547, plain, ((e13 = op1(e12, e13)) | ~ spl80_59), inference(avatar_component_clause, [], [f1545])).
fof(f1545, plain, (spl80_59 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_59])])).
fof(f9154, plain, ((e13 = op1(op1(e12, e13), e10)) | ~ spl80_36), inference(forward_demodulation, [], [f292, f1451])).
fof(f292, plain, (e13 = op1(op1(e12, e13), op1(e13, e12))), inference(cnf_transformation, [], [f3])).
fof(f10139, plain, (spl80_32 | ~ spl80_94), inference(avatar_split_clause, [], [f10133, f1692, f1432])).
fof(f1432, plain, (spl80_32 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_32])])).
fof(f1692, plain, (spl80_94 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_94])])).
fof(f10133, plain, ((e11 = op1(e13, e13)) | ~ spl80_94), inference(backward_demodulation, [], [f281, f1694])).
fof(f1694, plain, ((e13 = op1(e11, e11)) | ~ spl80_94), inference(avatar_component_clause, [], [f1692])).
fof(f281, plain, (e11 = op1(op1(e11, e11), op1(e11, e11))), inference(cnf_transformation, [], [f3])).
fof(f10130, plain, (~ spl80_75 | ~ spl80_100), inference(avatar_split_clause, [], [f10124, f1717, f1612])).
fof(f1612, plain, (spl80_75 <=> (e14 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_75])])).
fof(f10124, plain, (~ (e14 = op1(e12, e10)) | ~ spl80_100), inference(backward_demodulation, [], [f404, f1719])).
fof(f404, plain, ~ (op1(e11, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f7])).
fof(f10112, plain, (spl80_107 | ~ spl80_119), inference(avatar_split_clause, [], [f10107, f1797, f1747])).
fof(f10107, plain, ((e11 = op1(e10, e13)) | ~ spl80_119), inference(backward_demodulation, [], [f646, f1799])).
fof(f646, plain, (e11 = op1(e10, op1(e10, e11))), inference(cnf_transformation, [], [f12])).
fof(f10111, plain, (~ spl80_104 | ~ spl80_119), inference(avatar_split_clause, [], [f10106, f1797, f1734])).
fof(f1734, plain, (spl80_104 <=> (e13 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_104])])).
fof(f10106, plain, (~ (e13 = op1(e10, e14)) | ~ spl80_119), inference(backward_demodulation, [], [f456, f1799])).
fof(f456, plain, ~ (op1(e10, e11) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f10095, plain, (~ spl80_131 | ~ spl80_135), inference(avatar_contradiction_clause, [], [f10094])).
fof(f10094, plain, ($false | (~ spl80_131 | ~ spl80_135)), inference(subsumption_resolution, [], [f10093, f613])).
fof(f613, plain, ~ (e20 = e24), inference(cnf_transformation, [], [f10])).
fof(f10093, plain, ((e20 = e24) | (~ spl80_131 | ~ spl80_135)), inference(forward_demodulation, [], [f1916, f1900])).
fof(f1900, plain, ((e20 = op2(e24, e23)) | ~ spl80_131), inference(avatar_component_clause, [], [f1898])).
fof(f1898, plain, (spl80_131 <=> (e20 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_131])])).
fof(f1916, plain, ((e24 = op2(e24, e23)) | ~ spl80_135), inference(avatar_component_clause, [], [f1914])).
fof(f1914, plain, (spl80_135 <=> (e24 = op2(e24, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_135])])).
fof(f10067, plain, (~ spl80_161 | ~ spl80_162), inference(avatar_contradiction_clause, [], [f10066])).
fof(f10066, plain, ($false | (~ spl80_161 | ~ spl80_162)), inference(subsumption_resolution, [], [f10065, f610])).
fof(f610, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f10])).
fof(f10065, plain, ((e20 = e21) | (~ spl80_161 | ~ spl80_162)), inference(backward_demodulation, [], [f2030, f2026])).
fof(f2026, plain, ((e20 = op2(e23, e22)) | ~ spl80_161), inference(avatar_component_clause, [], [f2024])).
fof(f2024, plain, (spl80_161 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_161])])).
fof(f2030, plain, ((e21 = op2(e23, e22)) | ~ spl80_162), inference(avatar_component_clause, [], [f2028])).
fof(f2028, plain, (spl80_162 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_162])])).
fof(f10044, plain, (~ spl80_416 | ~ spl80_188), inference(avatar_split_clause, [], [f10038, f2137, f3560])).
fof(f10038, plain, (~ (e22 = h10(e14)) | ~ spl80_188), inference(backward_demodulation, [], [f2570, f2139])).
fof(f2570, plain, ~ (op2(e22, e22) = h10(e14)), inference(backward_demodulation, [], [f524, f750])).
fof(f524, plain, ~ (op2(e21, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f8])).
fof(f10019, plain, (~ spl80_1102 | ~ spl80_215), inference(avatar_split_clause, [], [f10016, f2250, f8167])).
fof(f2250, plain, (spl80_215 <=> (e24 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_215])])).
fof(f10016, plain, (~ (e24 = h1(e14)) | ~ spl80_215), inference(backward_demodulation, [], [f2464, f2252])).
fof(f2252, plain, ((e24 = op2(e21, e22)) | ~ spl80_215), inference(avatar_component_clause, [], [f2250])).
fof(f2464, plain, ~ (op2(e21, e22) = h1(e14)), inference(backward_demodulation, [], [f561, f705])).
fof(f561, plain, ~ (op2(e21, e20) = op2(e21, e22)), inference(cnf_transformation, [], [f8])).
fof(f10013, plain, (~ spl80_223 | spl80_552), inference(avatar_contradiction_clause, [], [f10012])).
fof(f10012, plain, ($false | (~ spl80_223 | spl80_552)), inference(subsumption_resolution, [], [f10011, f4307])).
fof(f4307, plain, (~ (e22 = h1(e14)) | spl80_552), inference(avatar_component_clause, [], [f4305])).
fof(f4305, plain, (spl80_552 <=> (e22 = h1(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_552])])).
fof(f10011, plain, ((e22 = h1(e14)) | ~ spl80_223), inference(backward_demodulation, [], [f705, f2286])).
fof(f2286, plain, ((e22 = op2(e21, e20)) | ~ spl80_223), inference(avatar_component_clause, [], [f2284])).
fof(f2284, plain, (spl80_223 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_223])])).
fof(f10006, plain, (spl80_880 | ~ spl80_240), inference(avatar_split_clause, [], [f10004, f2355, f6615])).
fof(f10004, plain, ((e24 = h9(e14)) | ~ spl80_240), inference(backward_demodulation, [], [f745, f2357])).
fof(f9965, plain, (spl80_352 | ~ spl80_208), inference(avatar_split_clause, [], [f9964, f2221, f3218])).
fof(f2221, plain, (spl80_208 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_208])])).
fof(f9964, plain, ((e22 = h14(e14)) | ~ spl80_208), inference(forward_demodulation, [], [f770, f2223])).
fof(f2223, plain, ((e22 = op2(e21, e23)) | ~ spl80_208), inference(avatar_component_clause, [], [f2221])).
fof(f770, plain, (op2(e21, e23) = h14(e14)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ((op2(e21, e23) = h14(e14)) & (h14(e12) = op2(e21, op2(op2(e21, e23), e21))) & (h14(e10) = op2(op2(e21, e23), e21)) & (e21 = h14(e13)) & (e23 = h14(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax29)).
fof(f9949, plain, (spl80_296 | ~ spl80_201), inference(avatar_split_clause, [], [f9948, f2192, f2919])).
fof(f9948, plain, ((e20 = h18(e14)) | ~ spl80_201), inference(forward_demodulation, [], [f790, f2194])).
fof(f9945, plain, (spl80_524 | ~ spl80_183 | ~ spl80_464), inference(avatar_split_clause, [], [f9944, f3822, f2116, f4154])).
fof(f2116, plain, (spl80_183 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_183])])).
fof(f9944, plain, ((e22 = h3(e10)) | (~ spl80_183 | ~ spl80_464)), inference(forward_demodulation, [], [f9384, f2118])).
fof(f2118, plain, ((e22 = op2(e22, e23)) | ~ spl80_183), inference(avatar_component_clause, [], [f2116])).
fof(f9872, plain, (spl80_328 | ~ spl80_131), inference(avatar_split_clause, [], [f9871, f1898, f3089])).
fof(f9871, plain, ((e20 = h16(e14)) | ~ spl80_131), inference(forward_demodulation, [], [f780, f1900])).
fof(f9861, plain, (~ spl80_448 | ~ spl80_128), inference(avatar_split_clause, [], [f9860, f1885, f3732])).
fof(f3732, plain, (spl80_448 <=> (e22 = h8(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_448])])).
fof(f9860, plain, (~ (e22 = h8(e14)) | ~ spl80_128), inference(forward_demodulation, [], [f2552, f1887])).
fof(f1887, plain, ((e22 = op2(e24, e24)) | ~ spl80_128), inference(avatar_component_clause, [], [f1885])).
fof(f2552, plain, ~ (op2(e24, e24) = h8(e14)), inference(backward_demodulation, [], [f596, f740])).
fof(f740, plain, (op2(e24, e21) = h8(e14)), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ((op2(e24, e21) = h8(e14)) & (h8(e12) = op2(e24, op2(op2(e24, e21), e24))) & (h8(e10) = op2(op2(e24, e21), e24)) & (e24 = h8(e13)) & (e21 = h8(e11))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax23)).
fof(f596, plain, ~ (op2(e24, e21) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f9822, plain, (spl80_54 | ~ spl80_45 | ~ spl80_83), inference(avatar_split_clause, [], [f9821, f1646, f1486, f1524])).
fof(f9821, plain, ((e13 = op1(e12, e14)) | (~ spl80_45 | ~ spl80_83)), inference(forward_demodulation, [], [f9197, f1648])).
fof(f9197, plain, ((e13 = op1(op1(e11, e13), e14)) | ~ spl80_45), inference(forward_demodulation, [], [f291, f1488])).
fof(f291, plain, (e13 = op1(op1(e11, e13), op1(e13, e11))), inference(cnf_transformation, [], [f3])).
fof(f9820, plain, (spl80_12 | ~ spl80_45 | ~ spl80_83), inference(avatar_split_clause, [], [f9819, f1646, f1486, f1348])).
fof(f9819, plain, ((e11 = op1(e14, e12)) | (~ spl80_45 | ~ spl80_83)), inference(forward_demodulation, [], [f9198, f1648])).
fof(f9198, plain, ((e11 = op1(e14, op1(e11, e13))) | ~ spl80_45), inference(forward_demodulation, [], [f283, f1488])).
fof(f283, plain, (e11 = op1(op1(e13, e11), op1(e11, e13))), inference(cnf_transformation, [], [f3])).
fof(f9796, plain, (~ spl80_70 | ~ spl80_45), inference(avatar_split_clause, [], [f9795, f1486, f1591])).
fof(f1591, plain, (spl80_70 <=> (e14 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_70])])).
fof(f9795, plain, (~ (e14 = op1(e12, e11)) | ~ spl80_45), inference(forward_demodulation, [], [f417, f1488])).
fof(f417, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f9794, plain, (~ spl80_68 | ~ spl80_18), inference(avatar_split_clause, [], [f9793, f1373, f1583])).
fof(f1583, plain, (spl80_68 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_68])])).
fof(f9793, plain, (~ (e12 = op1(e12, e11)) | ~ spl80_18), inference(forward_demodulation, [], [f418, f1375])).
fof(f418, plain, ~ (op1(e12, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f9786, plain, (spl80_53 | ~ spl80_65), inference(avatar_split_clause, [], [f9692, f1570, f1520])).
fof(f1520, plain, (spl80_53 <=> (e12 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_53])])).
fof(f1570, plain, (spl80_65 <=> (e14 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_65])])).
fof(f9692, plain, ((e12 = op1(e12, e14)) | ~ spl80_65), inference(backward_demodulation, [], [f657, f1572])).
fof(f1572, plain, ((e14 = op1(e12, e12)) | ~ spl80_65), inference(avatar_component_clause, [], [f1570])).
fof(f657, plain, (e12 = op1(e12, op1(e12, e12))), inference(cnf_transformation, [], [f12])).
fof(f9778, plain, (spl80_14 | ~ spl80_48 | ~ spl80_110), inference(avatar_split_clause, [], [f9606, f1759, f1499, f1356])).
fof(f1759, plain, (spl80_110 <=> (e14 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_110])])).
fof(f9606, plain, ((e13 = op1(e14, e12)) | (~ spl80_48 | ~ spl80_110)), inference(backward_demodulation, [], [f9249, f1761])).
fof(f1761, plain, ((e14 = op1(e10, e13)) | ~ spl80_110), inference(avatar_component_clause, [], [f1759])).
fof(f9708, plain, (~ spl80_36 | ~ spl80_58 | ~ spl80_72), inference(avatar_contradiction_clause, [], [f9707])).
fof(f9707, plain, ($false | (~ spl80_36 | ~ spl80_58 | ~ spl80_72)), inference(subsumption_resolution, [], [f9706, f605])).
fof(f605, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f9])).
fof(f9706, plain, ((e11 = e13) | (~ spl80_36 | ~ spl80_58 | ~ spl80_72)), inference(forward_demodulation, [], [f9702, f1602])).
fof(f9702, plain, ((e13 = op1(e12, e10)) | (~ spl80_36 | ~ spl80_58)), inference(backward_demodulation, [], [f9154, f1543])).
fof(f1543, plain, ((e12 = op1(e12, e13)) | ~ spl80_58), inference(avatar_component_clause, [], [f1541])).
fof(f1541, plain, (spl80_58 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_58])])).
fof(f9534, plain, (~ spl80_260 | ~ spl80_157), inference(avatar_split_clause, [], [f9525, f2007, f2728])).
fof(f9525, plain, (~ (e21 = h20(e14)) | ~ spl80_157), inference(backward_demodulation, [], [f2673, f2009])).
fof(f2009, plain, ((e21 = op2(e23, e23)) | ~ spl80_157), inference(avatar_component_clause, [], [f2007])).
fof(f2673, plain, ~ (op2(e23, e23) = h20(e14)), inference(backward_demodulation, [], [f589, f800])).
fof(f589, plain, ~ (op2(e23, e23) = op2(e23, e24)), inference(cnf_transformation, [], [f8])).
fof(f9494, plain, (spl80_1064 | ~ spl80_185 | ~ spl80_464), inference(avatar_split_clause, [], [f9492, f3822, f2124, f7917])).
fof(f9492, plain, ((e24 = h3(e10)) | (~ spl80_185 | ~ spl80_464)), inference(backward_demodulation, [], [f9384, f2126])).
fof(f2126, plain, ((e24 = op2(e22, e23)) | ~ spl80_185), inference(avatar_component_clause, [], [f2124])).
fof(f9484, plain, (~ spl80_540 | ~ spl80_187), inference(avatar_split_clause, [], [f9470, f2133, f4241])).
fof(f4241, plain, (spl80_540 <=> (e21 = h2(e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_540])])).
fof(f9470, plain, (~ (e21 = h2(e14)) | ~ spl80_187), inference(backward_demodulation, [], [f2477, f2135])).
fof(f2135, plain, ((e21 = op2(e22, e22)) | ~ spl80_187), inference(avatar_component_clause, [], [f2133])).
fof(f2477, plain, ~ (op2(e22, e22) = h2(e14)), inference(backward_demodulation, [], [f571, f710])).
fof(f571, plain, ~ (op2(e22, e20) = op2(e22, e22)), inference(cnf_transformation, [], [f8])).
fof(f9460, plain, (spl80_540 | ~ spl80_197), inference(avatar_split_clause, [], [f9458, f2175, f4241])).
fof(f9458, plain, ((e21 = h2(e14)) | ~ spl80_197), inference(backward_demodulation, [], [f710, f2177])).
fof(f9439, plain, (spl80_420 | ~ spl80_212), inference(avatar_split_clause, [], [f9436, f2238, f3582])).
fof(f2238, plain, (spl80_212 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_212])])).
fof(f9436, plain, ((e21 = h10(e14)) | ~ spl80_212), inference(backward_demodulation, [], [f750, f2240])).
fof(f2240, plain, ((e21 = op2(e21, e22)) | ~ spl80_212), inference(avatar_component_clause, [], [f2238])).
fof(f9406, plain, (spl80_436 | ~ spl80_237), inference(avatar_split_clause, [], [f9404, f2343, f3668])).
fof(f2343, plain, (spl80_237 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_237])])).
fof(f9404, plain, ((e21 = h9(e14)) | ~ spl80_237), inference(backward_demodulation, [], [f745, f2345])).
fof(f2345, plain, ((e21 = op2(e20, e22)) | ~ spl80_237), inference(avatar_component_clause, [], [f2343])).
fof(f9391, plain, (spl80_161 | ~ spl80_464), inference(avatar_split_clause, [], [f9383, f3822, f2024])).
fof(f9383, plain, ((e20 = op2(e23, e22)) | ~ spl80_464), inference(backward_demodulation, [], [f2492, f3823])).
fof(f2492, plain, (e20 = op2(e23, h3(e14))), inference(backward_demodulation, [], [f685, f715])).
fof(f685, plain, (e20 = op2(e23, op2(e23, e20))), inference(cnf_transformation, [], [f13])).
fof(f9390, plain, (~ spl80_158 | ~ spl80_464), inference(avatar_split_clause, [], [f9382, f3822, f2011])).
fof(f9382, plain, (~ (e22 = op2(e23, e23)) | ~ spl80_464), inference(backward_demodulation, [], [f2490, f3823])).
fof(f2490, plain, ~ (op2(e23, e23) = h3(e14)), inference(backward_demodulation, [], [f582, f715])).
fof(f582, plain, ~ (op2(e23, e20) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f9379, plain, (spl80_131 | ~ spl80_928), inference(avatar_split_clause, [], [f9372, f7002, f1898])).
fof(f9372, plain, ((e20 = op2(e24, e23)) | ~ spl80_928), inference(backward_demodulation, [], [f2541, f7003])).
fof(f9378, plain, (~ spl80_160 | ~ spl80_928), inference(avatar_split_clause, [], [f9370, f7002, f2019])).
fof(f2019, plain, (spl80_160 <=> (e24 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl80_160])])).
fof(f9370, plain, (~ (e24 = op2(e23, e23)) | ~ spl80_928), inference(backward_demodulation, [], [f2538, f7003])).
fof(f2538, plain, ~ (op2(e23, e23) = h7(e14)), inference(backward_demodulation, [], [f585, f735])).
fof(f585, plain, ~ (op2(e23, e21) = op2(e23, e23)), inference(cnf_transformation, [], [f8])).
fof(f9377, plain, (~ spl80_165 | ~ spl80_928), inference(avatar_split_clause, [], [f9369, f7002, f2040])).
fof(f2040, plain, (spl80_165 <=> (e24 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_165])])).
fof(f9369, plain, (~ (e24 = op2(e23, e22)) | ~ spl80_928), inference(backward_demodulation, [], [f2537, f7003])).
fof(f2537, plain, ~ (op2(e23, e22) = h7(e14)), inference(backward_demodulation, [], [f584, f735])).
fof(f584, plain, ~ (op2(e23, e21) = op2(e23, e22)), inference(cnf_transformation, [], [f8])).
fof(f9376, plain, (~ spl80_220 | ~ spl80_928), inference(avatar_split_clause, [], [f9368, f7002, f2271])).
fof(f2271, plain, (spl80_220 <=> (e24 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_220])])).
fof(f9368, plain, (~ (e24 = op2(e21, e21)) | ~ spl80_928), inference(backward_demodulation, [], [f2535, f7003])).
fof(f2535, plain, ~ (op2(e21, e21) = h7(e14)), inference(backward_demodulation, [], [f515, f735])).
fof(f515, plain, ~ (op2(e21, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f8])).
fof(f9328, plain, (spl80_260 | ~ spl80_152), inference(avatar_split_clause, [], [f9327, f1986, f2728])).
fof(f9327, plain, ((e21 = h20(e14)) | ~ spl80_152), inference(forward_demodulation, [], [f800, f1988])).
fof(f1988, plain, ((e21 = op2(e23, e24)) | ~ spl80_152), inference(avatar_component_clause, [], [f1986])).
fof(f9314, plain, (spl80_504 | ~ spl80_149), inference(avatar_split_clause, [], [f9313, f1973, f4043])).
fof(f9313, plain, ((e23 = h4(e14)) | ~ spl80_149), inference(forward_demodulation, [], [f720, f1975])).
fof(f9306, plain, (spl80_448 | ~ spl80_143), inference(avatar_split_clause, [], [f9305, f1948, f3732])).
fof(f9305, plain, ((e22 = h8(e14)) | ~ spl80_143), inference(forward_demodulation, [], [f740, f1950])).
fof(f9300, plain, (~ spl80_476 | ~ spl80_143), inference(avatar_split_clause, [], [f9299, f1948, f3892])).
fof(f9299, plain, (~ (e22 = h6(e14)) | ~ spl80_143), inference(forward_demodulation, [], [f2525, f1950])).
fof(f2525, plain, ~ (op2(e24, e21) = h6(e14)), inference(backward_demodulation, [], [f518, f730])).
fof(f518, plain, ~ (op2(e22, e21) = op2(e24, e21)), inference(cnf_transformation, [], [f8])).
fof(f9296, plain, (spl80_388 | ~ spl80_137), inference(avatar_split_clause, [], [f9295, f1923, f3411])).
fof(f9295, plain, ((e21 = h12(e14)) | ~ spl80_137), inference(forward_demodulation, [], [f760, f1925])).
fof(f9262, plain, (~ spl80_120 | ~ spl80_45), inference(avatar_split_clause, [], [f9261, f1486, f1801])).
fof(f1801, plain, (spl80_120 <=> (e14 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_120])])).
fof(f9261, plain, (~ (e14 = op1(e10, e11)) | ~ spl80_45), inference(forward_demodulation, [], [f412, f1488])).
fof(f412, plain, ~ (op1(e10, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f9226, plain, (spl80_120 | ~ spl80_102), inference(avatar_split_clause, [], [f9225, f1726, f1801])).
fof(f1726, plain, (spl80_102 <=> (e11 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_102])])).
fof(f9225, plain, ((e14 = op1(e10, e11)) | ~ spl80_102), inference(forward_demodulation, [], [f649, f1728])).
fof(f1728, plain, ((e11 = op1(e10, e14)) | ~ spl80_102), inference(avatar_component_clause, [], [f1726])).
fof(f649, plain, (e14 = op1(e10, op1(e10, e14))), inference(cnf_transformation, [], [f12])).
fof(f9215, plain, (~ spl80_91 | ~ spl80_66), inference(avatar_split_clause, [], [f9214, f1575, f1680])).
fof(f1680, plain, (spl80_91 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_91])])).
fof(f9214, plain, (~ (e10 = op1(e11, e11)) | ~ spl80_66), inference(forward_demodulation, [], [f414, f1577])).
fof(f414, plain, ~ (op1(e11, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f7])).
fof(f9213, plain, (~ spl80_95 | ~ spl80_45), inference(avatar_split_clause, [], [f9212, f1486, f1696])).
fof(f1696, plain, (spl80_95 <=> (e14 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_95])])).
fof(f9212, plain, (~ (e14 = op1(e11, e11)) | ~ spl80_45), inference(forward_demodulation, [], [f415, f1488])).
fof(f415, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f7])).
fof(f9170, plain, (~ spl80_73 | ~ spl80_48), inference(avatar_split_clause, [], [f9169, f1499, f1604])).
fof(f1604, plain, (spl80_73 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_73])])).
fof(f9169, plain, (~ (e12 = op1(e12, e10)) | ~ spl80_48), inference(forward_demodulation, [], [f407, f1501])).
fof(f407, plain, ~ (op1(e12, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f7])).
fof(f9162, plain, (spl80_72 | ~ spl80_66), inference(avatar_split_clause, [], [f8856, f1575, f1600])).
fof(f8856, plain, ((e11 = op1(e12, e10)) | ~ spl80_66), inference(backward_demodulation, [], [f656, f1577])).
fof(f656, plain, (e11 = op1(e12, op1(e12, e11))), inference(cnf_transformation, [], [f12])).
fof(f9126, plain, (~ spl80_32 | ~ spl80_27), inference(avatar_split_clause, [], [f9125, f1411, f1432])).
fof(f9125, plain, (~ (e11 = op1(e13, e13)) | ~ spl80_27), inference(forward_demodulation, [], [f489, f1413])).
fof(f489, plain, ~ (op1(e13, e13) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f9124, plain, (spl80_119 | ~ spl80_6 | ~ spl80_27), inference(avatar_split_clause, [], [f9123, f1411, f1323, f1797])).
fof(f9123, plain, ((e13 = op1(e10, e11)) | (~ spl80_6 | ~ spl80_27)), inference(forward_demodulation, [], [f9122, f1325])).
fof(f9122, plain, ((e13 = op1(op1(e14, e13), e11)) | ~ spl80_27), inference(forward_demodulation, [], [f294, f1413])).
fof(f294, plain, (e13 = op1(op1(e14, e13), op1(e13, e14))), inference(cnf_transformation, [], [f3])).
fof(f9121, plain, (spl80_100 | ~ spl80_6 | ~ spl80_27), inference(avatar_split_clause, [], [f9120, f1411, f1323, f1717])).
fof(f9120, plain, ((e14 = op1(e11, e10)) | (~ spl80_6 | ~ spl80_27)), inference(forward_demodulation, [], [f9119, f1413])).
fof(f9119, plain, ((e14 = op1(op1(e13, e14), e10)) | ~ spl80_6), inference(forward_demodulation, [], [f298, f1325])).
fof(f298, plain, (e14 = op1(op1(e13, e14), op1(e14, e13))), inference(cnf_transformation, [], [f3])).
fof(f9088, plain, (spl80_24 | ~ spl80_6), inference(avatar_split_clause, [], [f9087, f1323, f1398])).
fof(f9087, plain, ((e13 = op1(e14, e10)) | ~ spl80_6), inference(forward_demodulation, [], [f668, f1325])).
fof(f668, plain, (e13 = op1(e14, op1(e14, e13))), inference(cnf_transformation, [], [f12])).
fof(f9086, plain, (~ spl80_6 | ~ spl80_9), inference(avatar_contradiction_clause, [], [f9085])).
fof(f9085, plain, ($false | (~ spl80_6 | ~ spl80_9)), inference(subsumption_resolution, [], [f9084, f602])).
fof(f602, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f9])).
fof(f9084, plain, ((e10 = e13) | (~ spl80_6 | ~ spl80_9)), inference(forward_demodulation, [], [f1337, f1325])).
fof(f1337, plain, ((e13 = op1(e14, e13)) | ~ spl80_9), inference(avatar_component_clause, [], [f1335])).
fof(f1335, plain, (spl80_9 <=> (e13 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_9])])).
fof(f9057, plain, (spl80_4 | ~ spl80_10), inference(avatar_split_clause, [], [f9053, f1339, f1314])).
fof(f1339, plain, (spl80_10 <=> (e14 = op1(e14, e13))), introduced(avatar_definition, [new_symbols(naming, [spl80_10])])).
fof(f9053, plain, ((e13 = op1(e14, e14)) | ~ spl80_10), inference(backward_demodulation, [], [f668, f1341])).
fof(f1341, plain, ((e14 = op1(e14, e13)) | ~ spl80_10), inference(avatar_component_clause, [], [f1339])).
fof(f9041, plain, (spl80_3 | ~ spl80_15), inference(avatar_split_clause, [], [f9037, f1360, f1310])).
fof(f1360, plain, (spl80_15 <=> (e14 = op1(e14, e12))), introduced(avatar_definition, [new_symbols(naming, [spl80_15])])).
fof(f9037, plain, ((e12 = op1(e14, e14)) | ~ spl80_15), inference(backward_demodulation, [], [f667, f1362])).
fof(f1362, plain, ((e14 = op1(e14, e12)) | ~ spl80_15), inference(avatar_component_clause, [], [f1360])).
fof(f9013, plain, (~ spl80_24 | ~ spl80_25), inference(avatar_contradiction_clause, [], [f9012])).
fof(f9012, plain, ($false | (~ spl80_24 | ~ spl80_25)), inference(subsumption_resolution, [], [f9011, f609])).
fof(f609, plain, ~ (e13 = e14), inference(cnf_transformation, [], [f9])).
fof(f9011, plain, ((e13 = e14) | (~ spl80_24 | ~ spl80_25)), inference(backward_demodulation, [], [f1404, f1400])).
fof(f1404, plain, ((e14 = op1(e14, e10)) | ~ spl80_25), inference(avatar_component_clause, [], [f1402])).
fof(f1402, plain, (spl80_25 <=> (e14 = op1(e14, e10))), introduced(avatar_definition, [new_symbols(naming, [spl80_25])])).
fof(f8964, plain, (~ spl80_26 | ~ spl80_36), inference(avatar_split_clause, [], [f8958, f1449, f1407])).
fof(f1407, plain, (spl80_26 <=> (e10 = op1(e13, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_26])])).
fof(f8958, plain, (~ (e10 = op1(e13, e14)) | ~ spl80_36), inference(backward_demodulation, [], [f488, f1451])).
fof(f488, plain, ~ (op1(e13, e12) = op1(e13, e14)), inference(cnf_transformation, [], [f7])).
fof(f8951, plain, (spl80_6 | ~ spl80_45), inference(avatar_split_clause, [], [f8943, f1486, f1323])).
fof(f8943, plain, ((e10 = op1(e14, e13)) | ~ spl80_45), inference(backward_demodulation, [], [f695, f1488])).
fof(f695, plain, (e10 = op1(op1(e13, e11), e13)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((e14 = op1(e13, e11)) & (e12 = op1(e13, op1(op1(e13, e11), e13))) & (e10 = op1(op1(e13, e11), e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax14)).
fof(f8950, plain, (spl80_27 | ~ spl80_45), inference(avatar_split_clause, [], [f8942, f1486, f1411])).
fof(f8942, plain, ((e11 = op1(e13, e14)) | ~ spl80_45), inference(backward_demodulation, [], [f661, f1488])).
fof(f661, plain, (e11 = op1(e13, op1(e13, e11))), inference(cnf_transformation, [], [f12])).
fof(f8946, plain, (~ spl80_20 | ~ spl80_45), inference(avatar_split_clause, [], [f8938, f1486, f1381])).
fof(f1381, plain, (spl80_20 <=> (e14 = op1(e14, e11))), introduced(avatar_definition, [new_symbols(naming, [spl80_20])])).
fof(f8938, plain, (~ (e14 = op1(e14, e11)) | ~ spl80_45), inference(backward_demodulation, [], [f419, f1488])).
fof(f419, plain, ~ (op1(e13, e11) = op1(e14, e11)), inference(cnf_transformation, [], [f7])).
fof(f8934, plain, (spl80_36 | ~ spl80_48), inference(avatar_split_clause, [], [f8926, f1499, f1449])).
fof(f8926, plain, ((e10 = op1(e13, e12)) | ~ spl80_48), inference(backward_demodulation, [], [f660, f1501])).
fof(f660, plain, (e10 = op1(e13, op1(e13, e10))), inference(cnf_transformation, [], [f12])).
fof(f8917, plain, (~ spl80_53 | ~ spl80_54), inference(avatar_contradiction_clause, [], [f8916])).
fof(f8916, plain, ($false | (~ spl80_53 | ~ spl80_54)), inference(subsumption_resolution, [], [f8915, f607])).
fof(f8915, plain, ((e12 = e13) | (~ spl80_53 | ~ spl80_54)), inference(backward_demodulation, [], [f1526, f1522])).
fof(f1522, plain, ((e12 = op1(e12, e14)) | ~ spl80_53), inference(avatar_component_clause, [], [f1520])).
fof(f8909, plain, ~ spl80_55, inference(avatar_contradiction_clause, [], [f8908])).
fof(f8908, plain, ($false | ~ spl80_55), inference(subsumption_resolution, [], [f8907, f608])).
fof(f608, plain, ~ (e12 = e14), inference(cnf_transformation, [], [f9])).
fof(f8907, plain, ((e12 = e14) | ~ spl80_55), inference(forward_demodulation, [], [f8903, f667])).
fof(f8903, plain, ((e14 = op1(e14, op1(e14, e12))) | ~ spl80_55), inference(backward_demodulation, [], [f297, f1530])).
fof(f1530, plain, ((e14 = op1(e12, e14)) | ~ spl80_55), inference(avatar_component_clause, [], [f1528])).
fof(f1528, plain, (spl80_55 <=> (e14 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_55])])).
fof(f297, plain, (e14 = op1(op1(e12, e14), op1(e14, e12))), inference(cnf_transformation, [], [f3])).
fof(f8864, plain, (~ spl80_51 | ~ spl80_66), inference(avatar_split_clause, [], [f8855, f1575, f1512])).
fof(f1512, plain, (spl80_51 <=> (e10 = op1(e12, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_51])])).
fof(f8855, plain, (~ (e10 = op1(e12, e14)) | ~ spl80_66), inference(backward_demodulation, [], [f476, f1577])).
fof(f476, plain, ~ (op1(e12, e11) = op1(e12, e14)), inference(cnf_transformation, [], [f7])).
fof(f8726, plain, ~ spl80_105, inference(avatar_contradiction_clause, [], [f8725])).
fof(f8725, plain, ($false | ~ spl80_105), inference(subsumption_resolution, [], [f8724, f603])).
fof(f603, plain, ~ (e10 = e14), inference(cnf_transformation, [], [f9])).
fof(f8724, plain, ((e10 = e14) | ~ spl80_105), inference(forward_demodulation, [], [f8718, f665])).
fof(f665, plain, (e10 = op1(e14, op1(e14, e10))), inference(cnf_transformation, [], [f12])).
fof(f8718, plain, ((e14 = op1(e14, op1(e14, e10))) | ~ spl80_105), inference(backward_demodulation, [], [f295, f1740])).
fof(f1740, plain, ((e14 = op1(e10, e14)) | ~ spl80_105), inference(avatar_component_clause, [], [f1738])).
fof(f1738, plain, (spl80_105 <=> (e14 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_105])])).
fof(f295, plain, (e14 = op1(op1(e10, e14), op1(e14, e10))), inference(cnf_transformation, [], [f3])).
fof(f8662, plain, (~ spl80_101 | ~ spl80_121), inference(avatar_split_clause, [], [f8637, f1806, f1722])).
fof(f1722, plain, (spl80_101 <=> (e10 = op1(e10, e14))), introduced(avatar_definition, [new_symbols(naming, [spl80_101])])).
fof(f8637, plain, (~ (e10 = op1(e10, e14)) | ~ spl80_121), inference(backward_demodulation, [], [f453, f1808])).
fof(f453, plain, ~ (op1(e10, e10) = op1(e10, e14)), inference(cnf_transformation, [], [f7])).
fof(f8624, plain, (~ spl80_328 | ~ spl80_126), inference(avatar_split_clause, [], [f8609, f1877, f3089])).
fof(f8609, plain, (~ (e20 = h16(e14)) | ~ spl80_126), inference(backward_demodulation, [], [f2637, f1879])).
fof(f1879, plain, ((e20 = op2(e24, e24)) | ~ spl80_126), inference(avatar_component_clause, [], [f1877])).
fof(f2637, plain, ~ (op2(e24, e24) = h16(e14)), inference(backward_demodulation, [], [f599, f780])).
fof(f599, plain, ~ (op2(e24, e23) = op2(e24, e24)), inference(cnf_transformation, [], [f8])).
fof(f8566, plain, (spl80_796 | ~ spl80_140), inference(avatar_split_clause, [], [f8560, f1935, f6033])).
fof(f1935, plain, (spl80_140 <=> (e24 = op2(e24, e22))), introduced(avatar_definition, [new_symbols(naming, [spl80_140])])).
fof(f8560, plain, ((e24 = h12(e14)) | ~ spl80_140), inference(backward_demodulation, [], [f760, f1937])).
fof(f1937, plain, ((e24 = op2(e24, e22)) | ~ spl80_140), inference(avatar_component_clause, [], [f1935])).
fof(f8540, plain, (~ spl80_149 | ~ spl80_150), inference(avatar_contradiction_clause, [], [f8539])).
fof(f8539, plain, ($false | (~ spl80_149 | ~ spl80_150)), inference(subsumption_resolution, [], [f8538, f619])).
fof(f619, plain, ~ (e23 = e24), inference(cnf_transformation, [], [f10])).
fof(f8538, plain, ((e23 = e24) | (~ spl80_149 | ~ spl80_150)), inference(backward_demodulation, [], [f1979, f1975])).
fof(f1979, plain, ((e24 = op2(e24, e20)) | ~ spl80_150), inference(avatar_component_clause, [], [f1977])).
fof(f1977, plain, (spl80_150 <=> (e24 = op2(e24, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_150])])).
fof(f8504, plain, (~ spl80_156 | ~ spl80_173), inference(avatar_contradiction_clause, [], [f8503])).
fof(f8503, plain, ($false | (~ spl80_156 | ~ spl80_173)), inference(subsumption_resolution, [], [f8502, f617])).
fof(f617, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f10])).
fof(f8502, plain, ((e22 = e23) | (~ spl80_156 | ~ spl80_173)), inference(forward_demodulation, [], [f8490, f2076])).
fof(f2076, plain, ((e22 = op2(e23, e20)) | ~ spl80_173), inference(avatar_component_clause, [], [f2074])).
fof(f2074, plain, (spl80_173 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl80_173])])).
fof(f8490, plain, ((e23 = op2(e23, e20)) | ~ spl80_156), inference(backward_demodulation, [], [f688, f2005])).
fof(f2005, plain, ((e20 = op2(e23, e23)) | ~ spl80_156), inference(avatar_component_clause, [], [f2003])).
fof(f688, plain, (e23 = op2(e23, op2(e23, e23))), inference(cnf_transformation, [], [f13])).
fof(f8475, plain, (spl80_928 | ~ spl80_170), inference(avatar_split_clause, [], [f8471, f2061, f7002])).
fof(f2061, plain, (spl80_170 <=> (e24 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl80_170])])).
fof(f8471, plain, ((e24 = h7(e14)) | ~ spl80_170), inference(backward_demodulation, [], [f735, f2063])).
fof(f2063, plain, ((e24 = op2(e23, e21)) | ~ spl80_170), inference(avatar_component_clause, [], [f2061])).
fof(f8470, plain, (~ spl80_536 | ~ spl80_173), inference(avatar_split_clause, [], [f8467, f2074, f4219])).
fof(f8467, plain, (~ (e22 = h2(e14)) | ~ spl80_173), inference(backward_demodulation, [], [f2474, f2076])).
fof(f2474, plain, ~ (op2(e23, e20) = h2(e14)), inference(backward_demodulation, [], [f507, f710])).
fof(f507, plain, ~ (op2(e22, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f8])).
fof(f8469, plain, (~ spl80_552 | ~ spl80_173), inference(avatar_split_clause, [], [f8466, f2074, f4305])).
fof(f8466, plain, (~ (e22 = h1(e14)) | ~ spl80_173), inference(backward_demodulation, [], [f2461, f2076])).
fof(f2461, plain, ~ (op2(e23, e20) = h1(e14)), inference(backward_demodulation, [], [f505, f705])).
fof(f505, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f8])).
fof(f8468, plain, (spl80_464 | ~ spl80_173), inference(avatar_split_clause, [], [f8465, f2074, f3822])).
fof(f8465, plain, ((e22 = h3(e14)) | ~ spl80_173), inference(backward_demodulation, [], [f715, f2076])).
fof(f8407, plain, (spl80_484 | ~ spl80_191), inference(avatar_split_clause, [], [f8404, f2150, f3936])).
fof(f8404, plain, ((e20 = h6(e14)) | ~ spl80_191), inference(backward_demodulation, [], [f730, f2152])).
fof(f7101, plain, (~ spl80_927 | spl80_468 | spl80_465 | spl80_462 | spl80_459 | ~ spl80_928 | ~ spl80_929 | ~ spl80_930 | ~ spl80_931 | ~ spl80_932 | ~ spl80_933 | ~ spl80_934 | ~ spl80_935 | ~ spl80_936 | ~ spl80_937 | ~ spl80_938 | ~ spl80_939 | ~ spl80_940 | ~ spl80_941 | ~ spl80_942 | ~ spl80_943 | ~ spl80_944 | ~ spl80_945 | ~ spl80_946 | ~ spl80_947 | ~ spl80_948 | ~ spl80_949 | ~ spl80_950 | ~ spl80_951 | ~ spl80_952), inference(avatar_split_clause, [], [f6996, f7098, f7094, f7090, f7086, f7082, f7078, f7074, f7070, f7066, f7062, f7058, f7054, f7050, f7046, f7042, f7038, f7034, f7030, f7026, f7022, f7018, f7014, f7010, f7006, f7002, f3792, f3811, f3829, f3848, f6998])).
fof(f3848, plain, (spl80_468 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl80_468])])).
fof(f3829, plain, (spl80_465 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl80_465])])).
fof(f3811, plain, (spl80_462 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl80_462])])).
fof(f3792, plain, (spl80_459 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl80_459])])).
fof(f6996, plain, (~ (op2(e20, e20) = h7(op1(e10, e10))) | ~ (h5(e14) = h7(op1(e10, e11))) | ~ (h7(op1(e10, e12)) = op2(e20, h3(e14))) | ~ (h13(e14) = h7(op1(e10, e13))) | ~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | ~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14)))), inference(forward_demodulation, [], [f6995, f2546])).
fof(f6995, plain, (~ (h5(e14) = h7(op1(e10, e11))) | ~ (h7(op1(e10, e12)) = op2(e20, h3(e14))) | ~ (h13(e14) = h7(op1(e10, e13))) | ~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | ~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6994, f725])).
fof(f6994, plain, (~ (op2(e20, e21) = h7(op1(e10, e11))) | ~ (h7(op1(e10, e12)) = op2(e20, h3(e14))) | ~ (h13(e14) = h7(op1(e10, e13))) | ~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | ~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6993, f2546])).
fof(f6993, plain, (~ (h7(op1(e10, e11)) = op2(h7(e10), e21)) | ~ (h7(op1(e10, e12)) = op2(e20, h3(e14))) | ~ (h13(e14) = h7(op1(e10, e13))) | ~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | ~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6992, f731])).
fof(f6992, plain, (~ (h7(op1(e10, e12)) = op2(e20, h3(e14))) | ~ (h13(e14) = h7(op1(e10, e13))) | ~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | ~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6991, f2546])).
fof(f6991, plain, (~ (h7(op1(e10, e12)) = op2(h7(e10), h3(e14))) | ~ (h13(e14) = h7(op1(e10, e13))) | ~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | ~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6990, f2544])).
fof(f6990, plain, (~ (h13(e14) = h7(op1(e10, e13))) | ~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | ~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6989, f765])).
fof(f6989, plain, (~ (op2(e20, e23) = h7(op1(e10, e13))) | ~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | ~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6988, f2546])).
fof(f6988, plain, (~ (h7(op1(e10, e13)) = op2(h7(e10), e23)) | ~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | ~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6987, f732])).
fof(f6987, plain, (~ (h7(op1(e10, e14)) = op2(e20, h7(e14))) | ~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6986, f2546])).
fof(f6986, plain, (~ (h1(e14) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6985, f705])).
fof(f6985, plain, (~ (op2(e21, e20) = h7(op1(e11, e10))) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6984, f731])).
fof(f6984, plain, (~ (h7(op1(e11, e10)) = op2(h7(e11), e20)) | ~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6983, f2546])).
fof(f6983, plain, (~ (op2(e21, e21) = h7(op1(e11, e11))) | ~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6982, f731])).
fof(f6982, plain, (~ (h7(op1(e11, e12)) = op2(e21, h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6981, f731])).
fof(f6981, plain, (~ (h7(op1(e11, e12)) = op2(h7(e11), h3(e14))) | ~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6980, f2544])).
fof(f6980, plain, (~ (h14(e14) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6979, f770])).
fof(f6979, plain, (~ (op2(e21, e23) = h7(op1(e11, e13))) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6978, f731])).
fof(f6978, plain, (~ (h7(op1(e11, e13)) = op2(h7(e11), e23)) | ~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6977, f732])).
fof(f6977, plain, (~ (h7(op1(e11, e14)) = op2(e21, h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6976, f731])).
fof(f6976, plain, (~ (h7(op1(e12, e10)) = op2(h3(e14), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6975, f2544])).
fof(f6975, plain, (~ (h7(op1(e12, e10)) = op2(h7(e12), e20)) | ~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6974, f2546])).
fof(f6974, plain, (~ (h7(op1(e12, e11)) = op2(h3(e14), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6973, f2544])).
fof(f6973, plain, (~ (h7(op1(e12, e11)) = op2(h7(e12), e21)) | ~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6972, f731])).
fof(f6972, plain, (~ (op2(h3(e14), h3(e14)) = h7(op1(e12, e12))) | ~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6971, f2544])).
fof(f6971, plain, (~ (h3(e10) = h7(op1(e12, e13))) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6970, f2494])).
fof(f6970, plain, (~ (h7(op1(e12, e13)) = op2(h3(e14), e23)) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6969, f2544])).
fof(f6969, plain, (~ (h7(op1(e12, e13)) = op2(h7(e12), e23)) | ~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6968, f732])).
fof(f6968, plain, (~ (h7(op1(e12, e14)) = op2(h3(e14), h7(e14))) | ~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6967, f2544])).
fof(f6967, plain, (~ (h3(e14) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6966, f715])).
fof(f6966, plain, (~ (op2(e23, e20) = h7(op1(e13, e10))) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6965, f732])).
fof(f6965, plain, (~ (h7(op1(e13, e10)) = op2(h7(e13), e20)) | ~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6964, f2546])).
fof(f6964, plain, (~ (h7(e14) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6963, f735])).
fof(f6963, plain, (~ (op2(e23, e21) = h7(op1(e13, e11))) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6962, f732])).
fof(f6962, plain, (~ (h7(op1(e13, e11)) = op2(h7(e13), e21)) | ~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6961, f731])).
fof(f6961, plain, (~ (e20 = h7(op1(e13, e12))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6960, f2492])).
fof(f6960, plain, (~ (h7(op1(e13, e12)) = op2(e23, h3(e14))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6959, f732])).
fof(f6959, plain, (~ (h7(op1(e13, e12)) = op2(h7(e13), h3(e14))) | ~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6958, f2544])).
fof(f6958, plain, (~ (op2(e23, e23) = h7(op1(e13, e13))) | ~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6957, f732])).
fof(f6957, plain, (~ (e21 = h7(op1(e13, e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6956, f2540])).
fof(f6956, plain, (~ (h7(op1(e13, e14)) = op2(e23, h7(e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6955, f732])).
fof(f6955, plain, (~ (h7(op1(e14, e10)) = op2(h7(e14), e20)) | ~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6954, f2546])).
fof(f6954, plain, (~ (h7(op1(e14, e11)) = op2(h7(e14), e21)) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6953, f731])).
fof(f6953, plain, (~ (h7(op1(e14, e12)) = op2(h7(e14), h3(e14))) | ~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6952, f2544])).
fof(f6952, plain, (~ (e20 = h7(op1(e14, e13))) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f6951, f2541])).
fof(f6951, plain, (~ (h7(op1(e14, e13)) = op2(h7(e14), e23)) | ~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(forward_demodulation, [], [f1235, f732])).
fof(f1235, plain, (~ (e24 = h7(e14)) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))), inference(cnf_transformation, [], [f119])).
fof(f119, plain, (((~ (e24 = h20(e14)) & ~ (e24 = h20(e13)) & ~ (e24 = h20(e12)) & ~ (e24 = h20(e11)) & ~ (e24 = h20(e10))) | sP79 | sP78 | sP77 | sP76 | ~ (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) | ~ (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) | ~ (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) | ~ (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) | ~ (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) | ~ (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) | ~ (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) | ~ (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) | ~ (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) | ~ (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) | ~ (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) | ~ (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) | ~ (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) | ~ (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) | ~ (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) | ~ (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) | ~ (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) | ~ (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) | ~ (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) | ~ (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) | ~ (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) | ~ (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) | ~ (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) | ~ (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) | ~ (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) & ((~ (e24 = h19(e14)) & ~ (e24 = h19(e13)) & ~ (e24 = h19(e12)) & ~ (e24 = h19(e11)) & ~ (e24 = h19(e10))) | sP75 | sP74 | sP73 | sP72 | ~ (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) | ~ (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) | ~ (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) | ~ (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) | ~ (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) | ~ (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) | ~ (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) | ~ (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) | ~ (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) | ~ (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) | ~ (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) | ~ (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) | ~ (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) | ~ (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) | ~ (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) | ~ (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) | ~ (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) | ~ (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) | ~ (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) | ~ (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) | ~ (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) | ~ (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) | ~ (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) | ~ (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) | ~ (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) & ((~ (e24 = h18(e14)) & ~ (e24 = h18(e13)) & ~ (e24 = h18(e12)) & ~ (e24 = h18(e11)) & ~ (e24 = h18(e10))) | sP71 | sP70 | sP69 | sP68 | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) & ((~ (e24 = h17(e14)) & ~ (e24 = h17(e13)) & ~ (e24 = h17(e12)) & ~ (e24 = h17(e11)) & ~ (e24 = h17(e10))) | sP67 | sP66 | sP65 | sP64 | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) & ((~ (e24 = h16(e14)) & ~ (e24 = h16(e13)) & ~ (e24 = h16(e12)) & ~ (e24 = h16(e11)) & ~ (e24 = h16(e10))) | sP63 | sP62 | sP61 | sP60 | ~ (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) | ~ (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) | ~ (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) | ~ (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) | ~ (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) | ~ (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) | ~ (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) | ~ (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) | ~ (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) | ~ (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) | ~ (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) | ~ (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) | ~ (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) | ~ (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) | ~ (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) | ~ (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) | ~ (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) | ~ (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) | ~ (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) | ~ (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) | ~ (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) | ~ (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) | ~ (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) | ~ (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) | ~ (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) & ((~ (e24 = h15(e14)) & ~ (e24 = h15(e13)) & ~ (e24 = h15(e12)) & ~ (e24 = h15(e11)) & ~ (e24 = h15(e10))) | sP59 | sP58 | sP57 | sP56 | ~ (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) | ~ (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) | ~ (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) | ~ (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) | ~ (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) | ~ (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) | ~ (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) | ~ (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) | ~ (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) | ~ (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) | ~ (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) | ~ (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) | ~ (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) | ~ (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) | ~ (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) | ~ (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) | ~ (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) | ~ (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) | ~ (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) | ~ (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) | ~ (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) | ~ (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) | ~ (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) | ~ (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) | ~ (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) & ((~ (e24 = h14(e14)) & ~ (e24 = h14(e13)) & ~ (e24 = h14(e12)) & ~ (e24 = h14(e11)) & ~ (e24 = h14(e10))) | sP55 | sP54 | sP53 | sP52 | ~ (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) | ~ (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) | ~ (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) | ~ (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) | ~ (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) | ~ (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) | ~ (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) | ~ (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) | ~ (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) | ~ (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) | ~ (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) | ~ (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) | ~ (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) | ~ (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) | ~ (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) | ~ (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) | ~ (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) | ~ (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) | ~ (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) | ~ (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) | ~ (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) | ~ (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) | ~ (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) | ~ (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) | ~ (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) & ((~ (e24 = h13(e14)) & ~ (e24 = h13(e13)) & ~ (e24 = h13(e12)) & ~ (e24 = h13(e11)) & ~ (e24 = h13(e10))) | sP51 | sP50 | sP49 | sP48 | ~ (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) | ~ (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) | ~ (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) | ~ (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) | ~ (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) | ~ (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) | ~ (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) | ~ (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) | ~ (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) | ~ (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) | ~ (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) | ~ (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) | ~ (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) | ~ (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) | ~ (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) | ~ (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) | ~ (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) | ~ (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) | ~ (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) | ~ (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) | ~ (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) | ~ (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) | ~ (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) | ~ (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) | ~ (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) & ((~ (e24 = h12(e14)) & ~ (e24 = h12(e13)) & ~ (e24 = h12(e12)) & ~ (e24 = h12(e11)) & ~ (e24 = h12(e10))) | sP47 | sP46 | sP45 | sP44 | ~ (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) | ~ (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) | ~ (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) | ~ (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) | ~ (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) | ~ (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e24 = h11(e14)) & ~ (e24 = h11(e13)) & ~ (e24 = h11(e12)) & ~ (e24 = h11(e11)) & ~ (e24 = h11(e10))) | sP43 | sP42 | sP41 | sP40 | ~ (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) | ~ (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) | ~ (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) | ~ (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) | ~ (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) | ~ (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e24 = h10(e14)) & ~ (e24 = h10(e13)) & ~ (e24 = h10(e12)) & ~ (e24 = h10(e11)) & ~ (e24 = h10(e10))) | sP39 | sP38 | sP37 | sP36 | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e24 = h9(e14)) & ~ (e24 = h9(e13)) & ~ (e24 = h9(e12)) & ~ (e24 = h9(e11)) & ~ (e24 = h9(e10))) | sP35 | sP34 | sP33 | sP32 | ~ (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) | ~ (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) | ~ (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) | ~ (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) | ~ (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) | ~ (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e24 = h8(e14)) & ~ (e24 = h8(e13)) & ~ (e24 = h8(e12)) & ~ (e24 = h8(e11)) & ~ (e24 = h8(e10))) | sP31 | sP30 | sP29 | sP28 | ~ (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) | ~ (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) | ~ (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) | ~ (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) | ~ (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) | ~ (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | sP27 | sP26 | sP25 | sP24 | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | sP23 | sP22 | sP21 | sP20 | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | sP19 | sP18 | sP17 | sP16 | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | sP15 | sP14 | sP13 | sP12 | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | sP11 | sP10 | sP9 | sP8 | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | sP7 | sP6 | sP5 | sP4 | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | sP3 | sP2 | sP1 | sP0 | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f38, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39])).
fof(f39, plain, ((~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(usedef, [], [e39])).
fof(e39, plain, (sP0 <=> (~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f40, plain, ((~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(usedef, [], [e40])).
fof(e40, plain, (sP1 <=> (~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f41, plain, ((~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(usedef, [], [e41])).
fof(e41, plain, (sP2 <=> (~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f42, plain, ((~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | ~ sP3), inference(usedef, [], [e42])).
fof(e42, plain, (sP3 <=> (~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f43, plain, ((~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP4), inference(usedef, [], [e43])).
fof(e43, plain, (sP4 <=> (~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f44, plain, ((~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP5), inference(usedef, [], [e44])).
fof(e44, plain, (sP5 <=> (~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f45, plain, ((~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP6), inference(usedef, [], [e45])).
fof(e45, plain, (sP6 <=> (~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f46, plain, ((~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | ~ sP7), inference(usedef, [], [e46])).
fof(e46, plain, (sP7 <=> (~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f47, plain, ((~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP8), inference(usedef, [], [e47])).
fof(e47, plain, (sP8 <=> (~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f48, plain, ((~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP9), inference(usedef, [], [e48])).
fof(e48, plain, (sP9 <=> (~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f49, plain, ((~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP10), inference(usedef, [], [e49])).
fof(e49, plain, (sP10 <=> (~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f50, plain, ((~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | ~ sP11), inference(usedef, [], [e50])).
fof(e50, plain, (sP11 <=> (~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f51, plain, ((~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP12), inference(usedef, [], [e51])).
fof(e51, plain, (sP12 <=> (~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f52, plain, ((~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP13), inference(usedef, [], [e52])).
fof(e52, plain, (sP13 <=> (~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f53, plain, ((~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP14), inference(usedef, [], [e53])).
fof(e53, plain, (sP14 <=> (~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f54, plain, ((~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | ~ sP15), inference(usedef, [], [e54])).
fof(e54, plain, (sP15 <=> (~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f55, plain, ((~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP16), inference(usedef, [], [e55])).
fof(e55, plain, (sP16 <=> (~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f56, plain, ((~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP17), inference(usedef, [], [e56])).
fof(e56, plain, (sP17 <=> (~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f57, plain, ((~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP18), inference(usedef, [], [e57])).
fof(e57, plain, (sP18 <=> (~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f58, plain, ((~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | ~ sP19), inference(usedef, [], [e58])).
fof(e58, plain, (sP19 <=> (~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f59, plain, ((~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP20), inference(usedef, [], [e59])).
fof(e59, plain, (sP20 <=> (~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f60, plain, ((~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP21), inference(usedef, [], [e60])).
fof(e60, plain, (sP21 <=> (~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f61, plain, ((~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP22), inference(usedef, [], [e61])).
fof(e61, plain, (sP22 <=> (~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f62, plain, ((~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | ~ sP23), inference(usedef, [], [e62])).
fof(e62, plain, (sP23 <=> (~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f63, plain, ((~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP24), inference(usedef, [], [e63])).
fof(e63, plain, (sP24 <=> (~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f64, plain, ((~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP25), inference(usedef, [], [e64])).
fof(e64, plain, (sP25 <=> (~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f65, plain, ((~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP26), inference(usedef, [], [e65])).
fof(e65, plain, (sP26 <=> (~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f66, plain, ((~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | ~ sP27), inference(usedef, [], [e66])).
fof(e66, plain, (sP27 <=> (~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f67, plain, ((~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ sP28), inference(usedef, [], [e67])).
fof(e67, plain, (sP28 <=> (~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f68, plain, ((~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | ~ sP29), inference(usedef, [], [e68])).
fof(e68, plain, (sP29 <=> (~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f69, plain, ((~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | ~ sP30), inference(usedef, [], [e69])).
fof(e69, plain, (sP30 <=> (~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f70, plain, ((~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | ~ sP31), inference(usedef, [], [e70])).
fof(e70, plain, (sP31 <=> (~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f71, plain, ((~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ sP32), inference(usedef, [], [e71])).
fof(e71, plain, (sP32 <=> (~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f72, plain, ((~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | ~ sP33), inference(usedef, [], [e72])).
fof(e72, plain, (sP33 <=> (~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f73, plain, ((~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | ~ sP34), inference(usedef, [], [e73])).
fof(e73, plain, (sP34 <=> (~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f74, plain, ((~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | ~ sP35), inference(usedef, [], [e74])).
fof(e74, plain, (sP35 <=> (~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f75, plain, ((~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ sP36), inference(usedef, [], [e75])).
fof(e75, plain, (sP36 <=> (~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f76, plain, ((~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | ~ sP37), inference(usedef, [], [e76])).
fof(e76, plain, (sP37 <=> (~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f77, plain, ((~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | ~ sP38), inference(usedef, [], [e77])).
fof(e77, plain, (sP38 <=> (~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f78, plain, ((~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | ~ sP39), inference(usedef, [], [e78])).
fof(e78, plain, (sP39 <=> (~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f79, plain, ((~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ sP40), inference(usedef, [], [e79])).
fof(e79, plain, (sP40 <=> (~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f80, plain, ((~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | ~ sP41), inference(usedef, [], [e80])).
fof(e80, plain, (sP41 <=> (~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f81, plain, ((~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | ~ sP42), inference(usedef, [], [e81])).
fof(e81, plain, (sP42 <=> (~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f82, plain, ((~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | ~ sP43), inference(usedef, [], [e82])).
fof(e82, plain, (sP43 <=> (~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f83, plain, ((~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ sP44), inference(usedef, [], [e83])).
fof(e83, plain, (sP44 <=> (~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f84, plain, ((~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | ~ sP45), inference(usedef, [], [e84])).
fof(e84, plain, (sP45 <=> (~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f85, plain, ((~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | ~ sP46), inference(usedef, [], [e85])).
fof(e85, plain, (sP46 <=> (~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f86, plain, ((~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | ~ sP47), inference(usedef, [], [e86])).
fof(e86, plain, (sP47 <=> (~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f87, plain, ((~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10))) | ~ sP48), inference(usedef, [], [e87])).
fof(e87, plain, (sP48 <=> (~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f88, plain, ((~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10))) | ~ sP49), inference(usedef, [], [e88])).
fof(e88, plain, (sP49 <=> (~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f89, plain, ((~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10))) | ~ sP50), inference(usedef, [], [e89])).
fof(e89, plain, (sP50 <=> (~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f90, plain, ((~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10))) | ~ sP51), inference(usedef, [], [e90])).
fof(e90, plain, (sP51 <=> (~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f91, plain, ((~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10))) | ~ sP52), inference(usedef, [], [e91])).
fof(e91, plain, (sP52 <=> (~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f92, plain, ((~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10))) | ~ sP53), inference(usedef, [], [e92])).
fof(e92, plain, (sP53 <=> (~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f93, plain, ((~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10))) | ~ sP54), inference(usedef, [], [e93])).
fof(e93, plain, (sP54 <=> (~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f94, plain, ((~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10))) | ~ sP55), inference(usedef, [], [e94])).
fof(e94, plain, (sP55 <=> (~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f95, plain, ((~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10))) | ~ sP56), inference(usedef, [], [e95])).
fof(e95, plain, (sP56 <=> (~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f96, plain, ((~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10))) | ~ sP57), inference(usedef, [], [e96])).
fof(e96, plain, (sP57 <=> (~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f97, plain, ((~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10))) | ~ sP58), inference(usedef, [], [e97])).
fof(e97, plain, (sP58 <=> (~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f98, plain, ((~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10))) | ~ sP59), inference(usedef, [], [e98])).
fof(e98, plain, (sP59 <=> (~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f99, plain, ((~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10))) | ~ sP60), inference(usedef, [], [e99])).
fof(e99, plain, (sP60 <=> (~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f100, plain, ((~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10))) | ~ sP61), inference(usedef, [], [e100])).
fof(e100, plain, (sP61 <=> (~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f101, plain, ((~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10))) | ~ sP62), inference(usedef, [], [e101])).
fof(e101, plain, (sP62 <=> (~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f102, plain, ((~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10))) | ~ sP63), inference(usedef, [], [e102])).
fof(e102, plain, (sP63 <=> (~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f103, plain, ((~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10))) | ~ sP64), inference(usedef, [], [e103])).
fof(e103, plain, (sP64 <=> (~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f104, plain, ((~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10))) | ~ sP65), inference(usedef, [], [e104])).
fof(e104, plain, (sP65 <=> (~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f105, plain, ((~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10))) | ~ sP66), inference(usedef, [], [e105])).
fof(e105, plain, (sP66 <=> (~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f106, plain, ((~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10))) | ~ sP67), inference(usedef, [], [e106])).
fof(e106, plain, (sP67 <=> (~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f107, plain, ((~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10))) | ~ sP68), inference(usedef, [], [e107])).
fof(e107, plain, (sP68 <=> (~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f108, plain, ((~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10))) | ~ sP69), inference(usedef, [], [e108])).
fof(e108, plain, (sP69 <=> (~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f109, plain, ((~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10))) | ~ sP70), inference(usedef, [], [e109])).
fof(e109, plain, (sP70 <=> (~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f110, plain, ((~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10))) | ~ sP71), inference(usedef, [], [e110])).
fof(e110, plain, (sP71 <=> (~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f111, plain, ((~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10))) | ~ sP72), inference(usedef, [], [e111])).
fof(e111, plain, (sP72 <=> (~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f112, plain, ((~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10))) | ~ sP73), inference(usedef, [], [e112])).
fof(e112, plain, (sP73 <=> (~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f113, plain, ((~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10))) | ~ sP74), inference(usedef, [], [e113])).
fof(e113, plain, (sP74 <=> (~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f114, plain, ((~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10))) | ~ sP75), inference(usedef, [], [e114])).
fof(e114, plain, (sP75 <=> (~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f115, plain, ((~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10))) | ~ sP76), inference(usedef, [], [e115])).
fof(e115, plain, (sP76 <=> (~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f116, plain, ((~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10))) | ~ sP77), inference(usedef, [], [e116])).
fof(e116, plain, (sP77 <=> (~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f117, plain, ((~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10))) | ~ sP78), inference(usedef, [], [e117])).
fof(e117, plain, (sP78 <=> (~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f118, plain, ((~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10))) | ~ sP79), inference(usedef, [], [e118])).
fof(e118, plain, (sP79 <=> (~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f38, plain, (((~ (e24 = h20(e14)) & ~ (e24 = h20(e13)) & ~ (e24 = h20(e12)) & ~ (e24 = h20(e11)) & ~ (e24 = h20(e10))) | (~ (e23 = h20(e14)) & ~ (e23 = h20(e13)) & ~ (e23 = h20(e12)) & ~ (e23 = h20(e11)) & ~ (e23 = h20(e10))) | (~ (e22 = h20(e14)) & ~ (e22 = h20(e13)) & ~ (e22 = h20(e12)) & ~ (e22 = h20(e11)) & ~ (e22 = h20(e10))) | (~ (e21 = h20(e14)) & ~ (e21 = h20(e13)) & ~ (e21 = h20(e12)) & ~ (e21 = h20(e11)) & ~ (e21 = h20(e10))) | (~ (e20 = h20(e14)) & ~ (e20 = h20(e13)) & ~ (e20 = h20(e12)) & ~ (e20 = h20(e11)) & ~ (e20 = h20(e10))) | ~ (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) | ~ (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) | ~ (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) | ~ (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) | ~ (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) | ~ (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) | ~ (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) | ~ (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) | ~ (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) | ~ (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) | ~ (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) | ~ (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) | ~ (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) | ~ (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) | ~ (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) | ~ (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) | ~ (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) | ~ (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) | ~ (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) | ~ (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) | ~ (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) | ~ (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) | ~ (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) | ~ (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) | ~ (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) & ((~ (e24 = h19(e14)) & ~ (e24 = h19(e13)) & ~ (e24 = h19(e12)) & ~ (e24 = h19(e11)) & ~ (e24 = h19(e10))) | (~ (e23 = h19(e14)) & ~ (e23 = h19(e13)) & ~ (e23 = h19(e12)) & ~ (e23 = h19(e11)) & ~ (e23 = h19(e10))) | (~ (e22 = h19(e14)) & ~ (e22 = h19(e13)) & ~ (e22 = h19(e12)) & ~ (e22 = h19(e11)) & ~ (e22 = h19(e10))) | (~ (e21 = h19(e14)) & ~ (e21 = h19(e13)) & ~ (e21 = h19(e12)) & ~ (e21 = h19(e11)) & ~ (e21 = h19(e10))) | (~ (e20 = h19(e14)) & ~ (e20 = h19(e13)) & ~ (e20 = h19(e12)) & ~ (e20 = h19(e11)) & ~ (e20 = h19(e10))) | ~ (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) | ~ (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) | ~ (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) | ~ (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) | ~ (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) | ~ (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) | ~ (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) | ~ (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) | ~ (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) | ~ (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) | ~ (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) | ~ (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) | ~ (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) | ~ (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) | ~ (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) | ~ (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) | ~ (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) | ~ (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) | ~ (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) | ~ (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) | ~ (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) | ~ (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) | ~ (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) | ~ (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) | ~ (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) & ((~ (e24 = h18(e14)) & ~ (e24 = h18(e13)) & ~ (e24 = h18(e12)) & ~ (e24 = h18(e11)) & ~ (e24 = h18(e10))) | (~ (e23 = h18(e14)) & ~ (e23 = h18(e13)) & ~ (e23 = h18(e12)) & ~ (e23 = h18(e11)) & ~ (e23 = h18(e10))) | (~ (e22 = h18(e14)) & ~ (e22 = h18(e13)) & ~ (e22 = h18(e12)) & ~ (e22 = h18(e11)) & ~ (e22 = h18(e10))) | (~ (e21 = h18(e14)) & ~ (e21 = h18(e13)) & ~ (e21 = h18(e12)) & ~ (e21 = h18(e11)) & ~ (e21 = h18(e10))) | (~ (e20 = h18(e14)) & ~ (e20 = h18(e13)) & ~ (e20 = h18(e12)) & ~ (e20 = h18(e11)) & ~ (e20 = h18(e10))) | ~ (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) | ~ (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) | ~ (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) | ~ (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) | ~ (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) | ~ (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) | ~ (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) | ~ (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) | ~ (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) | ~ (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) | ~ (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) | ~ (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) | ~ (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) | ~ (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) | ~ (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) | ~ (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) | ~ (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) | ~ (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) | ~ (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) | ~ (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) | ~ (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) | ~ (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) | ~ (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) | ~ (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) | ~ (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) & ((~ (e24 = h17(e14)) & ~ (e24 = h17(e13)) & ~ (e24 = h17(e12)) & ~ (e24 = h17(e11)) & ~ (e24 = h17(e10))) | (~ (e23 = h17(e14)) & ~ (e23 = h17(e13)) & ~ (e23 = h17(e12)) & ~ (e23 = h17(e11)) & ~ (e23 = h17(e10))) | (~ (e22 = h17(e14)) & ~ (e22 = h17(e13)) & ~ (e22 = h17(e12)) & ~ (e22 = h17(e11)) & ~ (e22 = h17(e10))) | (~ (e21 = h17(e14)) & ~ (e21 = h17(e13)) & ~ (e21 = h17(e12)) & ~ (e21 = h17(e11)) & ~ (e21 = h17(e10))) | (~ (e20 = h17(e14)) & ~ (e20 = h17(e13)) & ~ (e20 = h17(e12)) & ~ (e20 = h17(e11)) & ~ (e20 = h17(e10))) | ~ (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) | ~ (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) | ~ (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) | ~ (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) | ~ (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) | ~ (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) | ~ (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) | ~ (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) | ~ (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) | ~ (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) | ~ (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) | ~ (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) | ~ (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) | ~ (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) | ~ (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) | ~ (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) | ~ (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) | ~ (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) | ~ (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) | ~ (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) | ~ (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) | ~ (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) | ~ (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) | ~ (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) | ~ (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) & ((~ (e24 = h16(e14)) & ~ (e24 = h16(e13)) & ~ (e24 = h16(e12)) & ~ (e24 = h16(e11)) & ~ (e24 = h16(e10))) | (~ (e23 = h16(e14)) & ~ (e23 = h16(e13)) & ~ (e23 = h16(e12)) & ~ (e23 = h16(e11)) & ~ (e23 = h16(e10))) | (~ (e22 = h16(e14)) & ~ (e22 = h16(e13)) & ~ (e22 = h16(e12)) & ~ (e22 = h16(e11)) & ~ (e22 = h16(e10))) | (~ (e21 = h16(e14)) & ~ (e21 = h16(e13)) & ~ (e21 = h16(e12)) & ~ (e21 = h16(e11)) & ~ (e21 = h16(e10))) | (~ (e20 = h16(e14)) & ~ (e20 = h16(e13)) & ~ (e20 = h16(e12)) & ~ (e20 = h16(e11)) & ~ (e20 = h16(e10))) | ~ (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) | ~ (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) | ~ (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) | ~ (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) | ~ (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) | ~ (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) | ~ (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) | ~ (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) | ~ (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) | ~ (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) | ~ (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) | ~ (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) | ~ (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) | ~ (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) | ~ (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) | ~ (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) | ~ (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) | ~ (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) | ~ (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) | ~ (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) | ~ (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) | ~ (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) | ~ (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) | ~ (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) | ~ (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) & ((~ (e24 = h15(e14)) & ~ (e24 = h15(e13)) & ~ (e24 = h15(e12)) & ~ (e24 = h15(e11)) & ~ (e24 = h15(e10))) | (~ (e23 = h15(e14)) & ~ (e23 = h15(e13)) & ~ (e23 = h15(e12)) & ~ (e23 = h15(e11)) & ~ (e23 = h15(e10))) | (~ (e22 = h15(e14)) & ~ (e22 = h15(e13)) & ~ (e22 = h15(e12)) & ~ (e22 = h15(e11)) & ~ (e22 = h15(e10))) | (~ (e21 = h15(e14)) & ~ (e21 = h15(e13)) & ~ (e21 = h15(e12)) & ~ (e21 = h15(e11)) & ~ (e21 = h15(e10))) | (~ (e20 = h15(e14)) & ~ (e20 = h15(e13)) & ~ (e20 = h15(e12)) & ~ (e20 = h15(e11)) & ~ (e20 = h15(e10))) | ~ (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) | ~ (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) | ~ (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) | ~ (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) | ~ (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) | ~ (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) | ~ (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) | ~ (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) | ~ (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) | ~ (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) | ~ (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) | ~ (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) | ~ (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) | ~ (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) | ~ (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) | ~ (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) | ~ (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) | ~ (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) | ~ (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) | ~ (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) | ~ (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) | ~ (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) | ~ (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) | ~ (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) | ~ (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) & ((~ (e24 = h14(e14)) & ~ (e24 = h14(e13)) & ~ (e24 = h14(e12)) & ~ (e24 = h14(e11)) & ~ (e24 = h14(e10))) | (~ (e23 = h14(e14)) & ~ (e23 = h14(e13)) & ~ (e23 = h14(e12)) & ~ (e23 = h14(e11)) & ~ (e23 = h14(e10))) | (~ (e22 = h14(e14)) & ~ (e22 = h14(e13)) & ~ (e22 = h14(e12)) & ~ (e22 = h14(e11)) & ~ (e22 = h14(e10))) | (~ (e21 = h14(e14)) & ~ (e21 = h14(e13)) & ~ (e21 = h14(e12)) & ~ (e21 = h14(e11)) & ~ (e21 = h14(e10))) | (~ (e20 = h14(e14)) & ~ (e20 = h14(e13)) & ~ (e20 = h14(e12)) & ~ (e20 = h14(e11)) & ~ (e20 = h14(e10))) | ~ (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) | ~ (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) | ~ (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) | ~ (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) | ~ (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) | ~ (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) | ~ (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) | ~ (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) | ~ (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) | ~ (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) | ~ (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) | ~ (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) | ~ (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) | ~ (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) | ~ (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) | ~ (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) | ~ (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) | ~ (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) | ~ (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) | ~ (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) | ~ (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) | ~ (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) | ~ (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) | ~ (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) | ~ (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) & ((~ (e24 = h13(e14)) & ~ (e24 = h13(e13)) & ~ (e24 = h13(e12)) & ~ (e24 = h13(e11)) & ~ (e24 = h13(e10))) | (~ (e23 = h13(e14)) & ~ (e23 = h13(e13)) & ~ (e23 = h13(e12)) & ~ (e23 = h13(e11)) & ~ (e23 = h13(e10))) | (~ (e22 = h13(e14)) & ~ (e22 = h13(e13)) & ~ (e22 = h13(e12)) & ~ (e22 = h13(e11)) & ~ (e22 = h13(e10))) | (~ (e21 = h13(e14)) & ~ (e21 = h13(e13)) & ~ (e21 = h13(e12)) & ~ (e21 = h13(e11)) & ~ (e21 = h13(e10))) | (~ (e20 = h13(e14)) & ~ (e20 = h13(e13)) & ~ (e20 = h13(e12)) & ~ (e20 = h13(e11)) & ~ (e20 = h13(e10))) | ~ (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) | ~ (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) | ~ (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) | ~ (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) | ~ (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) | ~ (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) | ~ (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) | ~ (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) | ~ (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) | ~ (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) | ~ (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) | ~ (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) | ~ (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) | ~ (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) | ~ (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) | ~ (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) | ~ (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) | ~ (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) | ~ (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) | ~ (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) | ~ (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) | ~ (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) | ~ (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) | ~ (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) | ~ (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) & ((~ (e24 = h12(e14)) & ~ (e24 = h12(e13)) & ~ (e24 = h12(e12)) & ~ (e24 = h12(e11)) & ~ (e24 = h12(e10))) | (~ (e23 = h12(e14)) & ~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | (~ (e22 = h12(e14)) & ~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | (~ (e21 = h12(e14)) & ~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | (~ (e20 = h12(e14)) & ~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) | ~ (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) | ~ (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) | ~ (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) | ~ (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) | ~ (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e24 = h11(e14)) & ~ (e24 = h11(e13)) & ~ (e24 = h11(e12)) & ~ (e24 = h11(e11)) & ~ (e24 = h11(e10))) | (~ (e23 = h11(e14)) & ~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | (~ (e22 = h11(e14)) & ~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | (~ (e21 = h11(e14)) & ~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | (~ (e20 = h11(e14)) & ~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) | ~ (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) | ~ (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) | ~ (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) | ~ (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) | ~ (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e24 = h10(e14)) & ~ (e24 = h10(e13)) & ~ (e24 = h10(e12)) & ~ (e24 = h10(e11)) & ~ (e24 = h10(e10))) | (~ (e23 = h10(e14)) & ~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | (~ (e22 = h10(e14)) & ~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | (~ (e21 = h10(e14)) & ~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | (~ (e20 = h10(e14)) & ~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) | ~ (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) | ~ (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) | ~ (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) | ~ (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) | ~ (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e24 = h9(e14)) & ~ (e24 = h9(e13)) & ~ (e24 = h9(e12)) & ~ (e24 = h9(e11)) & ~ (e24 = h9(e10))) | (~ (e23 = h9(e14)) & ~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | (~ (e22 = h9(e14)) & ~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | (~ (e21 = h9(e14)) & ~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | (~ (e20 = h9(e14)) & ~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) | ~ (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) | ~ (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) | ~ (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) | ~ (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) | ~ (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e24 = h8(e14)) & ~ (e24 = h8(e13)) & ~ (e24 = h8(e12)) & ~ (e24 = h8(e11)) & ~ (e24 = h8(e10))) | (~ (e23 = h8(e14)) & ~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | (~ (e22 = h8(e14)) & ~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | (~ (e21 = h8(e14)) & ~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | (~ (e20 = h8(e14)) & ~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) | ~ (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) | ~ (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) | ~ (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) | ~ (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) | ~ (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e24 = h7(e14)) & ~ (e24 = h7(e13)) & ~ (e24 = h7(e12)) & ~ (e24 = h7(e11)) & ~ (e24 = h7(e10))) | (~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | (~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | (~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | (~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) | ~ (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) | ~ (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) | ~ (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) | ~ (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) | ~ (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e24 = h6(e14)) & ~ (e24 = h6(e13)) & ~ (e24 = h6(e12)) & ~ (e24 = h6(e11)) & ~ (e24 = h6(e10))) | (~ (e23 = h6(e14)) & ~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e14)) & ~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e14)) & ~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e14)) & ~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) | ~ (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) | ~ (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) | ~ (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) | ~ (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) | ~ (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e24 = h5(e14)) & ~ (e24 = h5(e13)) & ~ (e24 = h5(e12)) & ~ (e24 = h5(e11)) & ~ (e24 = h5(e10))) | (~ (e23 = h5(e14)) & ~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e14)) & ~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e14)) & ~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e14)) & ~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) | ~ (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) | ~ (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) | ~ (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) | ~ (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) | ~ (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e24 = h4(e14)) & ~ (e24 = h4(e13)) & ~ (e24 = h4(e12)) & ~ (e24 = h4(e11)) & ~ (e24 = h4(e10))) | (~ (e23 = h4(e14)) & ~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e14)) & ~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e14)) & ~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e14)) & ~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) | ~ (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) | ~ (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) | ~ (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) | ~ (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) | ~ (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e24 = h3(e14)) & ~ (e24 = h3(e13)) & ~ (e24 = h3(e12)) & ~ (e24 = h3(e11)) & ~ (e24 = h3(e10))) | (~ (e23 = h3(e14)) & ~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e14)) & ~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e14)) & ~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e14)) & ~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) | ~ (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) | ~ (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) | ~ (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) | ~ (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) | ~ (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e24 = h2(e14)) & ~ (e24 = h2(e13)) & ~ (e24 = h2(e12)) & ~ (e24 = h2(e11)) & ~ (e24 = h2(e10))) | (~ (e23 = h2(e14)) & ~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e14)) & ~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e14)) & ~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e14)) & ~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) | ~ (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) | ~ (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) | ~ (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) | ~ (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) | ~ (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e24 = h1(e14)) & ~ (e24 = h1(e13)) & ~ (e24 = h1(e12)) & ~ (e24 = h1(e11)) & ~ (e24 = h1(e10))) | (~ (e23 = h1(e14)) & ~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e14)) & ~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e14)) & ~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e14)) & ~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) | ~ (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) | ~ (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) | ~ (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) | ~ (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) | ~ (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ~ ((((e24 = h20(e14)) | (e24 = h20(e13)) | (e24 = h20(e12)) | (e24 = h20(e11)) | (e24 = h20(e10))) & ((e23 = h20(e14)) | (e23 = h20(e13)) | (e23 = h20(e12)) | (e23 = h20(e11)) | (e23 = h20(e10))) & ((e22 = h20(e14)) | (e22 = h20(e13)) | (e22 = h20(e12)) | (e22 = h20(e11)) | (e22 = h20(e10))) & ((e21 = h20(e14)) | (e21 = h20(e13)) | (e21 = h20(e12)) | (e21 = h20(e11)) | (e21 = h20(e10))) & ((e20 = h20(e14)) | (e20 = h20(e13)) | (e20 = h20(e12)) | (e20 = h20(e11)) | (e20 = h20(e10))) & (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) & (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) & (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) & (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) & (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) & (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) & (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) & (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) & (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) & (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) & (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) & (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) & (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) & (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) & (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) & (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) & (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) & (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) & (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) & (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) & (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) & (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) & (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) & (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) & (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) | (((e24 = h19(e14)) | (e24 = h19(e13)) | (e24 = h19(e12)) | (e24 = h19(e11)) | (e24 = h19(e10))) & ((e23 = h19(e14)) | (e23 = h19(e13)) | (e23 = h19(e12)) | (e23 = h19(e11)) | (e23 = h19(e10))) & ((e22 = h19(e14)) | (e22 = h19(e13)) | (e22 = h19(e12)) | (e22 = h19(e11)) | (e22 = h19(e10))) & ((e21 = h19(e14)) | (e21 = h19(e13)) | (e21 = h19(e12)) | (e21 = h19(e11)) | (e21 = h19(e10))) & ((e20 = h19(e14)) | (e20 = h19(e13)) | (e20 = h19(e12)) | (e20 = h19(e11)) | (e20 = h19(e10))) & (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) & (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) & (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) & (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) & (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) & (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) & (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) & (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) & (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) & (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) & (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) & (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) & (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) & (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) & (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) & (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) & (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) & (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) & (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) & (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) & (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) & (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) & (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) & (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) & (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) | (((e24 = h18(e14)) | (e24 = h18(e13)) | (e24 = h18(e12)) | (e24 = h18(e11)) | (e24 = h18(e10))) & ((e23 = h18(e14)) | (e23 = h18(e13)) | (e23 = h18(e12)) | (e23 = h18(e11)) | (e23 = h18(e10))) & ((e22 = h18(e14)) | (e22 = h18(e13)) | (e22 = h18(e12)) | (e22 = h18(e11)) | (e22 = h18(e10))) & ((e21 = h18(e14)) | (e21 = h18(e13)) | (e21 = h18(e12)) | (e21 = h18(e11)) | (e21 = h18(e10))) & ((e20 = h18(e14)) | (e20 = h18(e13)) | (e20 = h18(e12)) | (e20 = h18(e11)) | (e20 = h18(e10))) & (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) & (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) & (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) & (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) & (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) & (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) & (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) & (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) & (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) & (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) & (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) & (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) & (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) & (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) & (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) & (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) & (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) & (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) & (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) & (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) & (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) & (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) & (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) & (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) & (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) | (((e24 = h17(e14)) | (e24 = h17(e13)) | (e24 = h17(e12)) | (e24 = h17(e11)) | (e24 = h17(e10))) & ((e23 = h17(e14)) | (e23 = h17(e13)) | (e23 = h17(e12)) | (e23 = h17(e11)) | (e23 = h17(e10))) & ((e22 = h17(e14)) | (e22 = h17(e13)) | (e22 = h17(e12)) | (e22 = h17(e11)) | (e22 = h17(e10))) & ((e21 = h17(e14)) | (e21 = h17(e13)) | (e21 = h17(e12)) | (e21 = h17(e11)) | (e21 = h17(e10))) & ((e20 = h17(e14)) | (e20 = h17(e13)) | (e20 = h17(e12)) | (e20 = h17(e11)) | (e20 = h17(e10))) & (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) & (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) & (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) & (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) & (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) & (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) & (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) & (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) & (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) & (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) & (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) & (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) & (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) & (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) & (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) & (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) & (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) & (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) & (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) & (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) & (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) & (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) & (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) & (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) & (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) | (((e24 = h16(e14)) | (e24 = h16(e13)) | (e24 = h16(e12)) | (e24 = h16(e11)) | (e24 = h16(e10))) & ((e23 = h16(e14)) | (e23 = h16(e13)) | (e23 = h16(e12)) | (e23 = h16(e11)) | (e23 = h16(e10))) & ((e22 = h16(e14)) | (e22 = h16(e13)) | (e22 = h16(e12)) | (e22 = h16(e11)) | (e22 = h16(e10))) & ((e21 = h16(e14)) | (e21 = h16(e13)) | (e21 = h16(e12)) | (e21 = h16(e11)) | (e21 = h16(e10))) & ((e20 = h16(e14)) | (e20 = h16(e13)) | (e20 = h16(e12)) | (e20 = h16(e11)) | (e20 = h16(e10))) & (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) & (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) & (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) & (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) & (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) & (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) & (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) & (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) & (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) & (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) & (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) & (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) & (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) & (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) & (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) & (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) & (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) & (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) & (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) & (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) & (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) & (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) & (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) & (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) & (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) | (((e24 = h15(e14)) | (e24 = h15(e13)) | (e24 = h15(e12)) | (e24 = h15(e11)) | (e24 = h15(e10))) & ((e23 = h15(e14)) | (e23 = h15(e13)) | (e23 = h15(e12)) | (e23 = h15(e11)) | (e23 = h15(e10))) & ((e22 = h15(e14)) | (e22 = h15(e13)) | (e22 = h15(e12)) | (e22 = h15(e11)) | (e22 = h15(e10))) & ((e21 = h15(e14)) | (e21 = h15(e13)) | (e21 = h15(e12)) | (e21 = h15(e11)) | (e21 = h15(e10))) & ((e20 = h15(e14)) | (e20 = h15(e13)) | (e20 = h15(e12)) | (e20 = h15(e11)) | (e20 = h15(e10))) & (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) & (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) & (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) & (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) & (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) & (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) & (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) & (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) & (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) & (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) & (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) & (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) & (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) & (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) & (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) & (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) & (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) & (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) & (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) & (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) & (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) & (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) & (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) & (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) & (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) | (((e24 = h14(e14)) | (e24 = h14(e13)) | (e24 = h14(e12)) | (e24 = h14(e11)) | (e24 = h14(e10))) & ((e23 = h14(e14)) | (e23 = h14(e13)) | (e23 = h14(e12)) | (e23 = h14(e11)) | (e23 = h14(e10))) & ((e22 = h14(e14)) | (e22 = h14(e13)) | (e22 = h14(e12)) | (e22 = h14(e11)) | (e22 = h14(e10))) & ((e21 = h14(e14)) | (e21 = h14(e13)) | (e21 = h14(e12)) | (e21 = h14(e11)) | (e21 = h14(e10))) & ((e20 = h14(e14)) | (e20 = h14(e13)) | (e20 = h14(e12)) | (e20 = h14(e11)) | (e20 = h14(e10))) & (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) & (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) & (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) & (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) & (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) & (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) & (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) & (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) & (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) & (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) & (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) & (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) & (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) & (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) & (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) & (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) & (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) & (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) & (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) & (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) & (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) & (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) & (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) & (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) & (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) | (((e24 = h13(e14)) | (e24 = h13(e13)) | (e24 = h13(e12)) | (e24 = h13(e11)) | (e24 = h13(e10))) & ((e23 = h13(e14)) | (e23 = h13(e13)) | (e23 = h13(e12)) | (e23 = h13(e11)) | (e23 = h13(e10))) & ((e22 = h13(e14)) | (e22 = h13(e13)) | (e22 = h13(e12)) | (e22 = h13(e11)) | (e22 = h13(e10))) & ((e21 = h13(e14)) | (e21 = h13(e13)) | (e21 = h13(e12)) | (e21 = h13(e11)) | (e21 = h13(e10))) & ((e20 = h13(e14)) | (e20 = h13(e13)) | (e20 = h13(e12)) | (e20 = h13(e11)) | (e20 = h13(e10))) & (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) & (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) & (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) & (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) & (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) & (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) & (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) & (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) & (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) & (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) & (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) & (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) & (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) & (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) & (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) & (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) & (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) & (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) & (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) & (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) & (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) & (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) & (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) & (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) & (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) | (((e24 = h12(e14)) | (e24 = h12(e13)) | (e24 = h12(e12)) | (e24 = h12(e11)) | (e24 = h12(e10))) & ((e23 = h12(e14)) | (e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e14)) | (e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e14)) | (e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e14)) | (e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) & (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) & (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) & (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) & (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) & (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e24 = h11(e14)) | (e24 = h11(e13)) | (e24 = h11(e12)) | (e24 = h11(e11)) | (e24 = h11(e10))) & ((e23 = h11(e14)) | (e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e14)) | (e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e14)) | (e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e14)) | (e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) & (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) & (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) & (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) & (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) & (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e24 = h10(e14)) | (e24 = h10(e13)) | (e24 = h10(e12)) | (e24 = h10(e11)) | (e24 = h10(e10))) & ((e23 = h10(e14)) | (e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e14)) | (e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e14)) | (e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e14)) | (e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) & (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) & (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) & (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) & (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) & (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e24 = h9(e14)) | (e24 = h9(e13)) | (e24 = h9(e12)) | (e24 = h9(e11)) | (e24 = h9(e10))) & ((e23 = h9(e14)) | (e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e14)) | (e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e14)) | (e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e14)) | (e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) & (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) & (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) & (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) & (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) & (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e24 = h8(e14)) | (e24 = h8(e13)) | (e24 = h8(e12)) | (e24 = h8(e11)) | (e24 = h8(e10))) & ((e23 = h8(e14)) | (e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e14)) | (e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e14)) | (e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e14)) | (e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) & (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) & (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) & (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) & (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) & (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f36])).
fof(f36, plain, ~ ((((e24 = h20(e14)) | (e24 = h20(e13)) | (e24 = h20(e12)) | (e24 = h20(e11)) | (e24 = h20(e10))) & ((e23 = h20(e14)) | (e23 = h20(e13)) | (e23 = h20(e12)) | (e23 = h20(e11)) | (e23 = h20(e10))) & ((e22 = h20(e14)) | (e22 = h20(e13)) | (e22 = h20(e12)) | (e22 = h20(e11)) | (e22 = h20(e10))) & ((e21 = h20(e14)) | (e21 = h20(e13)) | (e21 = h20(e12)) | (e21 = h20(e11)) | (e21 = h20(e10))) & ((e20 = h20(e14)) | (e20 = h20(e13)) | (e20 = h20(e12)) | (e20 = h20(e11)) | (e20 = h20(e10))) & (h20(op1(e14, e14)) = op2(h20(e14), h20(e14))) & (h20(op1(e14, e13)) = op2(h20(e14), h20(e13))) & (h20(op1(e14, e12)) = op2(h20(e14), h20(e12))) & (h20(op1(e14, e11)) = op2(h20(e14), h20(e11))) & (h20(op1(e14, e10)) = op2(h20(e14), h20(e10))) & (h20(op1(e13, e14)) = op2(h20(e13), h20(e14))) & (h20(op1(e13, e13)) = op2(h20(e13), h20(e13))) & (h20(op1(e13, e12)) = op2(h20(e13), h20(e12))) & (h20(op1(e13, e11)) = op2(h20(e13), h20(e11))) & (h20(op1(e13, e10)) = op2(h20(e13), h20(e10))) & (h20(op1(e12, e14)) = op2(h20(e12), h20(e14))) & (h20(op1(e12, e13)) = op2(h20(e12), h20(e13))) & (h20(op1(e12, e12)) = op2(h20(e12), h20(e12))) & (h20(op1(e12, e11)) = op2(h20(e12), h20(e11))) & (h20(op1(e12, e10)) = op2(h20(e12), h20(e10))) & (h20(op1(e11, e14)) = op2(h20(e11), h20(e14))) & (h20(op1(e11, e13)) = op2(h20(e11), h20(e13))) & (h20(op1(e11, e12)) = op2(h20(e11), h20(e12))) & (h20(op1(e11, e11)) = op2(h20(e11), h20(e11))) & (h20(op1(e11, e10)) = op2(h20(e11), h20(e10))) & (h20(op1(e10, e14)) = op2(h20(e10), h20(e14))) & (h20(op1(e10, e13)) = op2(h20(e10), h20(e13))) & (h20(op1(e10, e12)) = op2(h20(e10), h20(e12))) & (h20(op1(e10, e11)) = op2(h20(e10), h20(e11))) & (h20(op1(e10, e10)) = op2(h20(e10), h20(e10)))) | (((e24 = h19(e14)) | (e24 = h19(e13)) | (e24 = h19(e12)) | (e24 = h19(e11)) | (e24 = h19(e10))) & ((e23 = h19(e14)) | (e23 = h19(e13)) | (e23 = h19(e12)) | (e23 = h19(e11)) | (e23 = h19(e10))) & ((e22 = h19(e14)) | (e22 = h19(e13)) | (e22 = h19(e12)) | (e22 = h19(e11)) | (e22 = h19(e10))) & ((e21 = h19(e14)) | (e21 = h19(e13)) | (e21 = h19(e12)) | (e21 = h19(e11)) | (e21 = h19(e10))) & ((e20 = h19(e14)) | (e20 = h19(e13)) | (e20 = h19(e12)) | (e20 = h19(e11)) | (e20 = h19(e10))) & (h19(op1(e14, e14)) = op2(h19(e14), h19(e14))) & (h19(op1(e14, e13)) = op2(h19(e14), h19(e13))) & (h19(op1(e14, e12)) = op2(h19(e14), h19(e12))) & (h19(op1(e14, e11)) = op2(h19(e14), h19(e11))) & (h19(op1(e14, e10)) = op2(h19(e14), h19(e10))) & (h19(op1(e13, e14)) = op2(h19(e13), h19(e14))) & (h19(op1(e13, e13)) = op2(h19(e13), h19(e13))) & (h19(op1(e13, e12)) = op2(h19(e13), h19(e12))) & (h19(op1(e13, e11)) = op2(h19(e13), h19(e11))) & (h19(op1(e13, e10)) = op2(h19(e13), h19(e10))) & (h19(op1(e12, e14)) = op2(h19(e12), h19(e14))) & (h19(op1(e12, e13)) = op2(h19(e12), h19(e13))) & (h19(op1(e12, e12)) = op2(h19(e12), h19(e12))) & (h19(op1(e12, e11)) = op2(h19(e12), h19(e11))) & (h19(op1(e12, e10)) = op2(h19(e12), h19(e10))) & (h19(op1(e11, e14)) = op2(h19(e11), h19(e14))) & (h19(op1(e11, e13)) = op2(h19(e11), h19(e13))) & (h19(op1(e11, e12)) = op2(h19(e11), h19(e12))) & (h19(op1(e11, e11)) = op2(h19(e11), h19(e11))) & (h19(op1(e11, e10)) = op2(h19(e11), h19(e10))) & (h19(op1(e10, e14)) = op2(h19(e10), h19(e14))) & (h19(op1(e10, e13)) = op2(h19(e10), h19(e13))) & (h19(op1(e10, e12)) = op2(h19(e10), h19(e12))) & (h19(op1(e10, e11)) = op2(h19(e10), h19(e11))) & (h19(op1(e10, e10)) = op2(h19(e10), h19(e10)))) | (((e24 = h18(e14)) | (e24 = h18(e13)) | (e24 = h18(e12)) | (e24 = h18(e11)) | (e24 = h18(e10))) & ((e23 = h18(e14)) | (e23 = h18(e13)) | (e23 = h18(e12)) | (e23 = h18(e11)) | (e23 = h18(e10))) & ((e22 = h18(e14)) | (e22 = h18(e13)) | (e22 = h18(e12)) | (e22 = h18(e11)) | (e22 = h18(e10))) & ((e21 = h18(e14)) | (e21 = h18(e13)) | (e21 = h18(e12)) | (e21 = h18(e11)) | (e21 = h18(e10))) & ((e20 = h18(e14)) | (e20 = h18(e13)) | (e20 = h18(e12)) | (e20 = h18(e11)) | (e20 = h18(e10))) & (h18(op1(e14, e14)) = op2(h18(e14), h18(e14))) & (h18(op1(e14, e13)) = op2(h18(e14), h18(e13))) & (h18(op1(e14, e12)) = op2(h18(e14), h18(e12))) & (h18(op1(e14, e11)) = op2(h18(e14), h18(e11))) & (h18(op1(e14, e10)) = op2(h18(e14), h18(e10))) & (h18(op1(e13, e14)) = op2(h18(e13), h18(e14))) & (h18(op1(e13, e13)) = op2(h18(e13), h18(e13))) & (h18(op1(e13, e12)) = op2(h18(e13), h18(e12))) & (h18(op1(e13, e11)) = op2(h18(e13), h18(e11))) & (h18(op1(e13, e10)) = op2(h18(e13), h18(e10))) & (h18(op1(e12, e14)) = op2(h18(e12), h18(e14))) & (h18(op1(e12, e13)) = op2(h18(e12), h18(e13))) & (h18(op1(e12, e12)) = op2(h18(e12), h18(e12))) & (h18(op1(e12, e11)) = op2(h18(e12), h18(e11))) & (h18(op1(e12, e10)) = op2(h18(e12), h18(e10))) & (h18(op1(e11, e14)) = op2(h18(e11), h18(e14))) & (h18(op1(e11, e13)) = op2(h18(e11), h18(e13))) & (h18(op1(e11, e12)) = op2(h18(e11), h18(e12))) & (h18(op1(e11, e11)) = op2(h18(e11), h18(e11))) & (h18(op1(e11, e10)) = op2(h18(e11), h18(e10))) & (h18(op1(e10, e14)) = op2(h18(e10), h18(e14))) & (h18(op1(e10, e13)) = op2(h18(e10), h18(e13))) & (h18(op1(e10, e12)) = op2(h18(e10), h18(e12))) & (h18(op1(e10, e11)) = op2(h18(e10), h18(e11))) & (h18(op1(e10, e10)) = op2(h18(e10), h18(e10)))) | (((e24 = h17(e14)) | (e24 = h17(e13)) | (e24 = h17(e12)) | (e24 = h17(e11)) | (e24 = h17(e10))) & ((e23 = h17(e14)) | (e23 = h17(e13)) | (e23 = h17(e12)) | (e23 = h17(e11)) | (e23 = h17(e10))) & ((e22 = h17(e14)) | (e22 = h17(e13)) | (e22 = h17(e12)) | (e22 = h17(e11)) | (e22 = h17(e10))) & ((e21 = h17(e14)) | (e21 = h17(e13)) | (e21 = h17(e12)) | (e21 = h17(e11)) | (e21 = h17(e10))) & ((e20 = h17(e14)) | (e20 = h17(e13)) | (e20 = h17(e12)) | (e20 = h17(e11)) | (e20 = h17(e10))) & (h17(op1(e14, e14)) = op2(h17(e14), h17(e14))) & (h17(op1(e14, e13)) = op2(h17(e14), h17(e13))) & (h17(op1(e14, e12)) = op2(h17(e14), h17(e12))) & (h17(op1(e14, e11)) = op2(h17(e14), h17(e11))) & (h17(op1(e14, e10)) = op2(h17(e14), h17(e10))) & (h17(op1(e13, e14)) = op2(h17(e13), h17(e14))) & (h17(op1(e13, e13)) = op2(h17(e13), h17(e13))) & (h17(op1(e13, e12)) = op2(h17(e13), h17(e12))) & (h17(op1(e13, e11)) = op2(h17(e13), h17(e11))) & (h17(op1(e13, e10)) = op2(h17(e13), h17(e10))) & (h17(op1(e12, e14)) = op2(h17(e12), h17(e14))) & (h17(op1(e12, e13)) = op2(h17(e12), h17(e13))) & (h17(op1(e12, e12)) = op2(h17(e12), h17(e12))) & (h17(op1(e12, e11)) = op2(h17(e12), h17(e11))) & (h17(op1(e12, e10)) = op2(h17(e12), h17(e10))) & (h17(op1(e11, e14)) = op2(h17(e11), h17(e14))) & (h17(op1(e11, e13)) = op2(h17(e11), h17(e13))) & (h17(op1(e11, e12)) = op2(h17(e11), h17(e12))) & (h17(op1(e11, e11)) = op2(h17(e11), h17(e11))) & (h17(op1(e11, e10)) = op2(h17(e11), h17(e10))) & (h17(op1(e10, e14)) = op2(h17(e10), h17(e14))) & (h17(op1(e10, e13)) = op2(h17(e10), h17(e13))) & (h17(op1(e10, e12)) = op2(h17(e10), h17(e12))) & (h17(op1(e10, e11)) = op2(h17(e10), h17(e11))) & (h17(op1(e10, e10)) = op2(h17(e10), h17(e10)))) | (((e24 = h16(e14)) | (e24 = h16(e13)) | (e24 = h16(e12)) | (e24 = h16(e11)) | (e24 = h16(e10))) & ((e23 = h16(e14)) | (e23 = h16(e13)) | (e23 = h16(e12)) | (e23 = h16(e11)) | (e23 = h16(e10))) & ((e22 = h16(e14)) | (e22 = h16(e13)) | (e22 = h16(e12)) | (e22 = h16(e11)) | (e22 = h16(e10))) & ((e21 = h16(e14)) | (e21 = h16(e13)) | (e21 = h16(e12)) | (e21 = h16(e11)) | (e21 = h16(e10))) & ((e20 = h16(e14)) | (e20 = h16(e13)) | (e20 = h16(e12)) | (e20 = h16(e11)) | (e20 = h16(e10))) & (h16(op1(e14, e14)) = op2(h16(e14), h16(e14))) & (h16(op1(e14, e13)) = op2(h16(e14), h16(e13))) & (h16(op1(e14, e12)) = op2(h16(e14), h16(e12))) & (h16(op1(e14, e11)) = op2(h16(e14), h16(e11))) & (h16(op1(e14, e10)) = op2(h16(e14), h16(e10))) & (h16(op1(e13, e14)) = op2(h16(e13), h16(e14))) & (h16(op1(e13, e13)) = op2(h16(e13), h16(e13))) & (h16(op1(e13, e12)) = op2(h16(e13), h16(e12))) & (h16(op1(e13, e11)) = op2(h16(e13), h16(e11))) & (h16(op1(e13, e10)) = op2(h16(e13), h16(e10))) & (h16(op1(e12, e14)) = op2(h16(e12), h16(e14))) & (h16(op1(e12, e13)) = op2(h16(e12), h16(e13))) & (h16(op1(e12, e12)) = op2(h16(e12), h16(e12))) & (h16(op1(e12, e11)) = op2(h16(e12), h16(e11))) & (h16(op1(e12, e10)) = op2(h16(e12), h16(e10))) & (h16(op1(e11, e14)) = op2(h16(e11), h16(e14))) & (h16(op1(e11, e13)) = op2(h16(e11), h16(e13))) & (h16(op1(e11, e12)) = op2(h16(e11), h16(e12))) & (h16(op1(e11, e11)) = op2(h16(e11), h16(e11))) & (h16(op1(e11, e10)) = op2(h16(e11), h16(e10))) & (h16(op1(e10, e14)) = op2(h16(e10), h16(e14))) & (h16(op1(e10, e13)) = op2(h16(e10), h16(e13))) & (h16(op1(e10, e12)) = op2(h16(e10), h16(e12))) & (h16(op1(e10, e11)) = op2(h16(e10), h16(e11))) & (h16(op1(e10, e10)) = op2(h16(e10), h16(e10)))) | (((e24 = h15(e14)) | (e24 = h15(e13)) | (e24 = h15(e12)) | (e24 = h15(e11)) | (e24 = h15(e10))) & ((e23 = h15(e14)) | (e23 = h15(e13)) | (e23 = h15(e12)) | (e23 = h15(e11)) | (e23 = h15(e10))) & ((e22 = h15(e14)) | (e22 = h15(e13)) | (e22 = h15(e12)) | (e22 = h15(e11)) | (e22 = h15(e10))) & ((e21 = h15(e14)) | (e21 = h15(e13)) | (e21 = h15(e12)) | (e21 = h15(e11)) | (e21 = h15(e10))) & ((e20 = h15(e14)) | (e20 = h15(e13)) | (e20 = h15(e12)) | (e20 = h15(e11)) | (e20 = h15(e10))) & (h15(op1(e14, e14)) = op2(h15(e14), h15(e14))) & (h15(op1(e14, e13)) = op2(h15(e14), h15(e13))) & (h15(op1(e14, e12)) = op2(h15(e14), h15(e12))) & (h15(op1(e14, e11)) = op2(h15(e14), h15(e11))) & (h15(op1(e14, e10)) = op2(h15(e14), h15(e10))) & (h15(op1(e13, e14)) = op2(h15(e13), h15(e14))) & (h15(op1(e13, e13)) = op2(h15(e13), h15(e13))) & (h15(op1(e13, e12)) = op2(h15(e13), h15(e12))) & (h15(op1(e13, e11)) = op2(h15(e13), h15(e11))) & (h15(op1(e13, e10)) = op2(h15(e13), h15(e10))) & (h15(op1(e12, e14)) = op2(h15(e12), h15(e14))) & (h15(op1(e12, e13)) = op2(h15(e12), h15(e13))) & (h15(op1(e12, e12)) = op2(h15(e12), h15(e12))) & (h15(op1(e12, e11)) = op2(h15(e12), h15(e11))) & (h15(op1(e12, e10)) = op2(h15(e12), h15(e10))) & (h15(op1(e11, e14)) = op2(h15(e11), h15(e14))) & (h15(op1(e11, e13)) = op2(h15(e11), h15(e13))) & (h15(op1(e11, e12)) = op2(h15(e11), h15(e12))) & (h15(op1(e11, e11)) = op2(h15(e11), h15(e11))) & (h15(op1(e11, e10)) = op2(h15(e11), h15(e10))) & (h15(op1(e10, e14)) = op2(h15(e10), h15(e14))) & (h15(op1(e10, e13)) = op2(h15(e10), h15(e13))) & (h15(op1(e10, e12)) = op2(h15(e10), h15(e12))) & (h15(op1(e10, e11)) = op2(h15(e10), h15(e11))) & (h15(op1(e10, e10)) = op2(h15(e10), h15(e10)))) | (((e24 = h14(e14)) | (e24 = h14(e13)) | (e24 = h14(e12)) | (e24 = h14(e11)) | (e24 = h14(e10))) & ((e23 = h14(e14)) | (e23 = h14(e13)) | (e23 = h14(e12)) | (e23 = h14(e11)) | (e23 = h14(e10))) & ((e22 = h14(e14)) | (e22 = h14(e13)) | (e22 = h14(e12)) | (e22 = h14(e11)) | (e22 = h14(e10))) & ((e21 = h14(e14)) | (e21 = h14(e13)) | (e21 = h14(e12)) | (e21 = h14(e11)) | (e21 = h14(e10))) & ((e20 = h14(e14)) | (e20 = h14(e13)) | (e20 = h14(e12)) | (e20 = h14(e11)) | (e20 = h14(e10))) & (h14(op1(e14, e14)) = op2(h14(e14), h14(e14))) & (h14(op1(e14, e13)) = op2(h14(e14), h14(e13))) & (h14(op1(e14, e12)) = op2(h14(e14), h14(e12))) & (h14(op1(e14, e11)) = op2(h14(e14), h14(e11))) & (h14(op1(e14, e10)) = op2(h14(e14), h14(e10))) & (h14(op1(e13, e14)) = op2(h14(e13), h14(e14))) & (h14(op1(e13, e13)) = op2(h14(e13), h14(e13))) & (h14(op1(e13, e12)) = op2(h14(e13), h14(e12))) & (h14(op1(e13, e11)) = op2(h14(e13), h14(e11))) & (h14(op1(e13, e10)) = op2(h14(e13), h14(e10))) & (h14(op1(e12, e14)) = op2(h14(e12), h14(e14))) & (h14(op1(e12, e13)) = op2(h14(e12), h14(e13))) & (h14(op1(e12, e12)) = op2(h14(e12), h14(e12))) & (h14(op1(e12, e11)) = op2(h14(e12), h14(e11))) & (h14(op1(e12, e10)) = op2(h14(e12), h14(e10))) & (h14(op1(e11, e14)) = op2(h14(e11), h14(e14))) & (h14(op1(e11, e13)) = op2(h14(e11), h14(e13))) & (h14(op1(e11, e12)) = op2(h14(e11), h14(e12))) & (h14(op1(e11, e11)) = op2(h14(e11), h14(e11))) & (h14(op1(e11, e10)) = op2(h14(e11), h14(e10))) & (h14(op1(e10, e14)) = op2(h14(e10), h14(e14))) & (h14(op1(e10, e13)) = op2(h14(e10), h14(e13))) & (h14(op1(e10, e12)) = op2(h14(e10), h14(e12))) & (h14(op1(e10, e11)) = op2(h14(e10), h14(e11))) & (h14(op1(e10, e10)) = op2(h14(e10), h14(e10)))) | (((e24 = h13(e14)) | (e24 = h13(e13)) | (e24 = h13(e12)) | (e24 = h13(e11)) | (e24 = h13(e10))) & ((e23 = h13(e14)) | (e23 = h13(e13)) | (e23 = h13(e12)) | (e23 = h13(e11)) | (e23 = h13(e10))) & ((e22 = h13(e14)) | (e22 = h13(e13)) | (e22 = h13(e12)) | (e22 = h13(e11)) | (e22 = h13(e10))) & ((e21 = h13(e14)) | (e21 = h13(e13)) | (e21 = h13(e12)) | (e21 = h13(e11)) | (e21 = h13(e10))) & ((e20 = h13(e14)) | (e20 = h13(e13)) | (e20 = h13(e12)) | (e20 = h13(e11)) | (e20 = h13(e10))) & (h13(op1(e14, e14)) = op2(h13(e14), h13(e14))) & (h13(op1(e14, e13)) = op2(h13(e14), h13(e13))) & (h13(op1(e14, e12)) = op2(h13(e14), h13(e12))) & (h13(op1(e14, e11)) = op2(h13(e14), h13(e11))) & (h13(op1(e14, e10)) = op2(h13(e14), h13(e10))) & (h13(op1(e13, e14)) = op2(h13(e13), h13(e14))) & (h13(op1(e13, e13)) = op2(h13(e13), h13(e13))) & (h13(op1(e13, e12)) = op2(h13(e13), h13(e12))) & (h13(op1(e13, e11)) = op2(h13(e13), h13(e11))) & (h13(op1(e13, e10)) = op2(h13(e13), h13(e10))) & (h13(op1(e12, e14)) = op2(h13(e12), h13(e14))) & (h13(op1(e12, e13)) = op2(h13(e12), h13(e13))) & (h13(op1(e12, e12)) = op2(h13(e12), h13(e12))) & (h13(op1(e12, e11)) = op2(h13(e12), h13(e11))) & (h13(op1(e12, e10)) = op2(h13(e12), h13(e10))) & (h13(op1(e11, e14)) = op2(h13(e11), h13(e14))) & (h13(op1(e11, e13)) = op2(h13(e11), h13(e13))) & (h13(op1(e11, e12)) = op2(h13(e11), h13(e12))) & (h13(op1(e11, e11)) = op2(h13(e11), h13(e11))) & (h13(op1(e11, e10)) = op2(h13(e11), h13(e10))) & (h13(op1(e10, e14)) = op2(h13(e10), h13(e14))) & (h13(op1(e10, e13)) = op2(h13(e10), h13(e13))) & (h13(op1(e10, e12)) = op2(h13(e10), h13(e12))) & (h13(op1(e10, e11)) = op2(h13(e10), h13(e11))) & (h13(op1(e10, e10)) = op2(h13(e10), h13(e10)))) | (((e24 = h12(e14)) | (e24 = h12(e13)) | (e24 = h12(e12)) | (e24 = h12(e11)) | (e24 = h12(e10))) & ((e23 = h12(e14)) | (e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e14)) | (e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e14)) | (e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e14)) | (e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e14, e14)) = op2(h12(e14), h12(e14))) & (h12(op1(e14, e13)) = op2(h12(e14), h12(e13))) & (h12(op1(e14, e12)) = op2(h12(e14), h12(e12))) & (h12(op1(e14, e11)) = op2(h12(e14), h12(e11))) & (h12(op1(e14, e10)) = op2(h12(e14), h12(e10))) & (h12(op1(e13, e14)) = op2(h12(e13), h12(e14))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e14)) = op2(h12(e12), h12(e14))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e14)) = op2(h12(e11), h12(e14))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e14)) = op2(h12(e10), h12(e14))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e24 = h11(e14)) | (e24 = h11(e13)) | (e24 = h11(e12)) | (e24 = h11(e11)) | (e24 = h11(e10))) & ((e23 = h11(e14)) | (e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e14)) | (e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e14)) | (e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e14)) | (e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e14, e14)) = op2(h11(e14), h11(e14))) & (h11(op1(e14, e13)) = op2(h11(e14), h11(e13))) & (h11(op1(e14, e12)) = op2(h11(e14), h11(e12))) & (h11(op1(e14, e11)) = op2(h11(e14), h11(e11))) & (h11(op1(e14, e10)) = op2(h11(e14), h11(e10))) & (h11(op1(e13, e14)) = op2(h11(e13), h11(e14))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e14)) = op2(h11(e12), h11(e14))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e14)) = op2(h11(e11), h11(e14))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e14)) = op2(h11(e10), h11(e14))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e24 = h10(e14)) | (e24 = h10(e13)) | (e24 = h10(e12)) | (e24 = h10(e11)) | (e24 = h10(e10))) & ((e23 = h10(e14)) | (e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e14)) | (e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e14)) | (e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e14)) | (e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e14, e14)) = op2(h10(e14), h10(e14))) & (h10(op1(e14, e13)) = op2(h10(e14), h10(e13))) & (h10(op1(e14, e12)) = op2(h10(e14), h10(e12))) & (h10(op1(e14, e11)) = op2(h10(e14), h10(e11))) & (h10(op1(e14, e10)) = op2(h10(e14), h10(e10))) & (h10(op1(e13, e14)) = op2(h10(e13), h10(e14))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e14)) = op2(h10(e12), h10(e14))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e14)) = op2(h10(e11), h10(e14))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e14)) = op2(h10(e10), h10(e14))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e24 = h9(e14)) | (e24 = h9(e13)) | (e24 = h9(e12)) | (e24 = h9(e11)) | (e24 = h9(e10))) & ((e23 = h9(e14)) | (e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e14)) | (e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e14)) | (e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e14)) | (e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e14, e14)) = op2(h9(e14), h9(e14))) & (h9(op1(e14, e13)) = op2(h9(e14), h9(e13))) & (h9(op1(e14, e12)) = op2(h9(e14), h9(e12))) & (h9(op1(e14, e11)) = op2(h9(e14), h9(e11))) & (h9(op1(e14, e10)) = op2(h9(e14), h9(e10))) & (h9(op1(e13, e14)) = op2(h9(e13), h9(e14))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e14)) = op2(h9(e12), h9(e14))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e14)) = op2(h9(e11), h9(e14))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e14)) = op2(h9(e10), h9(e14))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e24 = h8(e14)) | (e24 = h8(e13)) | (e24 = h8(e12)) | (e24 = h8(e11)) | (e24 = h8(e10))) & ((e23 = h8(e14)) | (e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e14)) | (e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e14)) | (e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e14)) | (e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e14, e14)) = op2(h8(e14), h8(e14))) & (h8(op1(e14, e13)) = op2(h8(e14), h8(e13))) & (h8(op1(e14, e12)) = op2(h8(e14), h8(e12))) & (h8(op1(e14, e11)) = op2(h8(e14), h8(e11))) & (h8(op1(e14, e10)) = op2(h8(e14), h8(e10))) & (h8(op1(e13, e14)) = op2(h8(e13), h8(e14))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e14)) = op2(h8(e12), h8(e14))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e14)) = op2(h8(e11), h8(e14))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e14)) = op2(h8(e10), h8(e14))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e24 = h7(e14)) | (e24 = h7(e13)) | (e24 = h7(e12)) | (e24 = h7(e11)) | (e24 = h7(e10))) & ((e23 = h7(e14)) | (e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e14)) | (e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e14)) | (e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e14)) | (e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e14, e14)) = op2(h7(e14), h7(e14))) & (h7(op1(e14, e13)) = op2(h7(e14), h7(e13))) & (h7(op1(e14, e12)) = op2(h7(e14), h7(e12))) & (h7(op1(e14, e11)) = op2(h7(e14), h7(e11))) & (h7(op1(e14, e10)) = op2(h7(e14), h7(e10))) & (h7(op1(e13, e14)) = op2(h7(e13), h7(e14))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e14)) = op2(h7(e12), h7(e14))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e14)) = op2(h7(e11), h7(e14))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e14)) = op2(h7(e10), h7(e14))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e24 = h6(e14)) | (e24 = h6(e13)) | (e24 = h6(e12)) | (e24 = h6(e11)) | (e24 = h6(e10))) & ((e23 = h6(e14)) | (e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e14)) | (e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e14)) | (e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e14)) | (e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e14, e14)) = op2(h6(e14), h6(e14))) & (h6(op1(e14, e13)) = op2(h6(e14), h6(e13))) & (h6(op1(e14, e12)) = op2(h6(e14), h6(e12))) & (h6(op1(e14, e11)) = op2(h6(e14), h6(e11))) & (h6(op1(e14, e10)) = op2(h6(e14), h6(e10))) & (h6(op1(e13, e14)) = op2(h6(e13), h6(e14))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e14)) = op2(h6(e12), h6(e14))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e14)) = op2(h6(e11), h6(e14))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e14)) = op2(h6(e10), h6(e14))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e24 = h5(e14)) | (e24 = h5(e13)) | (e24 = h5(e12)) | (e24 = h5(e11)) | (e24 = h5(e10))) & ((e23 = h5(e14)) | (e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e14)) | (e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e14)) | (e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e14)) | (e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e14, e14)) = op2(h5(e14), h5(e14))) & (h5(op1(e14, e13)) = op2(h5(e14), h5(e13))) & (h5(op1(e14, e12)) = op2(h5(e14), h5(e12))) & (h5(op1(e14, e11)) = op2(h5(e14), h5(e11))) & (h5(op1(e14, e10)) = op2(h5(e14), h5(e10))) & (h5(op1(e13, e14)) = op2(h5(e13), h5(e14))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e14)) = op2(h5(e12), h5(e14))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e14)) = op2(h5(e11), h5(e14))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e14)) = op2(h5(e10), h5(e14))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e24 = h4(e14)) | (e24 = h4(e13)) | (e24 = h4(e12)) | (e24 = h4(e11)) | (e24 = h4(e10))) & ((e23 = h4(e14)) | (e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e14)) | (e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e14)) | (e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e14)) | (e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e14, e14)) = op2(h4(e14), h4(e14))) & (h4(op1(e14, e13)) = op2(h4(e14), h4(e13))) & (h4(op1(e14, e12)) = op2(h4(e14), h4(e12))) & (h4(op1(e14, e11)) = op2(h4(e14), h4(e11))) & (h4(op1(e14, e10)) = op2(h4(e14), h4(e10))) & (h4(op1(e13, e14)) = op2(h4(e13), h4(e14))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e14)) = op2(h4(e12), h4(e14))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e14)) = op2(h4(e11), h4(e14))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e14)) = op2(h4(e10), h4(e14))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e24 = h3(e14)) | (e24 = h3(e13)) | (e24 = h3(e12)) | (e24 = h3(e11)) | (e24 = h3(e10))) & ((e23 = h3(e14)) | (e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e14)) | (e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e14)) | (e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e14)) | (e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e14, e14)) = op2(h3(e14), h3(e14))) & (h3(op1(e14, e13)) = op2(h3(e14), h3(e13))) & (h3(op1(e14, e12)) = op2(h3(e14), h3(e12))) & (h3(op1(e14, e11)) = op2(h3(e14), h3(e11))) & (h3(op1(e14, e10)) = op2(h3(e14), h3(e10))) & (h3(op1(e13, e14)) = op2(h3(e13), h3(e14))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e14)) = op2(h3(e12), h3(e14))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e14)) = op2(h3(e11), h3(e14))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e14)) = op2(h3(e10), h3(e14))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e24 = h2(e14)) | (e24 = h2(e13)) | (e24 = h2(e12)) | (e24 = h2(e11)) | (e24 = h2(e10))) & ((e23 = h2(e14)) | (e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e14)) | (e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e14)) | (e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e14)) | (e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e14, e14)) = op2(h2(e14), h2(e14))) & (h2(op1(e14, e13)) = op2(h2(e14), h2(e13))) & (h2(op1(e14, e12)) = op2(h2(e14), h2(e12))) & (h2(op1(e14, e11)) = op2(h2(e14), h2(e11))) & (h2(op1(e14, e10)) = op2(h2(e14), h2(e10))) & (h2(op1(e13, e14)) = op2(h2(e13), h2(e14))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e14)) = op2(h2(e12), h2(e14))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e14)) = op2(h2(e11), h2(e14))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e14)) = op2(h2(e10), h2(e14))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e24 = h1(e14)) | (e24 = h1(e13)) | (e24 = h1(e12)) | (e24 = h1(e11)) | (e24 = h1(e10))) & ((e23 = h1(e14)) | (e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e14)) | (e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e14)) | (e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e14)) | (e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e14, e14)) = op2(h1(e14), h1(e14))) & (h1(op1(e14, e13)) = op2(h1(e14), h1(e13))) & (h1(op1(e14, e12)) = op2(h1(e14), h1(e12))) & (h1(op1(e14, e11)) = op2(h1(e14), h1(e11))) & (h1(op1(e14, e10)) = op2(h1(e14), h1(e10))) & (h1(op1(e13, e14)) = op2(h1(e13), h1(e14))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e14)) = op2(h1(e12), h1(e14))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e14)) = op2(h1(e11), h1(e14))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e14)) = op2(h1(e10), h1(e14))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', co1)).
fof(f3865, plain, ~ spl80_468, inference(avatar_split_clause, [], [f3864, f3848])).
fof(f3864, plain, ~ sP24, inference(subsumption_resolution, [], [f1076, f2546])).
fof(f1076, plain, (~ (e20 = h7(e10)) | ~ sP24), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e20 = h7(e14)) & ~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP24), inference(nnf_transformation, [], [f63])).
fof(f3845, plain, ~ spl80_465, inference(avatar_split_clause, [], [f3844, f3829])).
fof(f3844, plain, ~ sP25, inference(subsumption_resolution, [], [f1072, f731])).
fof(f1072, plain, (~ (e21 = h7(e11)) | ~ sP25), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e21 = h7(e14)) & ~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP25), inference(nnf_transformation, [], [f64])).
fof(f3825, plain, (~ spl80_462 | ~ spl80_464), inference(avatar_split_clause, [], [f3820, f3822, f3811])).
fof(f3820, plain, (~ (e22 = h3(e14)) | ~ sP26), inference(forward_demodulation, [], [f1068, f2544])).
fof(f1068, plain, (~ (e22 = h7(e12)) | ~ sP26), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e22 = h7(e14)) & ~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP26), inference(nnf_transformation, [], [f65])).
fof(f3801, plain, ~ spl80_459, inference(avatar_split_clause, [], [f3800, f3792])).
fof(f3800, plain, ~ sP27, inference(subsumption_resolution, [], [f1064, f732])).
fof(f1064, plain, (~ (e23 = h7(e13)) | ~ sP27), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e23 = h7(e14)) & ~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | ~ sP27), inference(nnf_transformation, [], [f66])).
fof(f2456, plain, spl80_173, inference(avatar_split_clause, [], [f2455, f2074])).
fof(f2455, plain, (e22 = op2(e23, e20)), inference(backward_demodulation, [], [f699, f698])).
fof(f699, plain, (e22 = op2(e23, op2(op2(e23, e21), e23))), inference(cnf_transformation, [], [f15])).
fof(f2454, plain, spl80_170, inference(avatar_split_clause, [], [f700, f2061])).
fof(f700, plain, (e24 = op2(e23, e21)), inference(cnf_transformation, [], [f15])).
fof(f2453, plain, spl80_48, inference(avatar_split_clause, [], [f2452, f1499])).
fof(f2452, plain, (e12 = op1(e13, e10)), inference(backward_demodulation, [], [f696, f695])).
fof(f696, plain, (e12 = op1(e13, op1(op1(e13, e11), e13))), inference(cnf_transformation, [], [f14])).
fof(f2451, plain, spl80_45, inference(avatar_split_clause, [], [f697, f1486])).
fof(f697, plain, (e14 = op1(e13, e11)), inference(cnf_transformation, [], [f14])).
fof(f2436, plain, (spl80_223 | spl80_218 | spl80_213 | spl80_208 | spl80_203), inference(avatar_split_clause, [], [f339, f2200, f2221, f2242, f2263, f2284])).
fof(f339, plain, ((e22 = op2(e21, e24)) | (e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (((e24 = op2(e24, e24)) | (e24 = op2(e23, e24)) | (e24 = op2(e22, e24)) | (e24 = op2(e21, e24)) | (e24 = op2(e20, e24))) & ((e24 = op2(e24, e24)) | (e24 = op2(e24, e23)) | (e24 = op2(e24, e22)) | (e24 = op2(e24, e21)) | (e24 = op2(e24, e20))) & ((e23 = op2(e24, e24)) | (e23 = op2(e23, e24)) | (e23 = op2(e22, e24)) | (e23 = op2(e21, e24)) | (e23 = op2(e20, e24))) & ((e23 = op2(e24, e24)) | (e23 = op2(e24, e23)) | (e23 = op2(e24, e22)) | (e23 = op2(e24, e21)) | (e23 = op2(e24, e20))) & ((e22 = op2(e24, e24)) | (e22 = op2(e23, e24)) | (e22 = op2(e22, e24)) | (e22 = op2(e21, e24)) | (e22 = op2(e20, e24))) & ((e22 = op2(e24, e24)) | (e22 = op2(e24, e23)) | (e22 = op2(e24, e22)) | (e22 = op2(e24, e21)) | (e22 = op2(e24, e20))) & ((e21 = op2(e24, e24)) | (e21 = op2(e23, e24)) | (e21 = op2(e22, e24)) | (e21 = op2(e21, e24)) | (e21 = op2(e20, e24))) & ((e21 = op2(e24, e24)) | (e21 = op2(e24, e23)) | (e21 = op2(e24, e22)) | (e21 = op2(e24, e21)) | (e21 = op2(e24, e20))) & ((e20 = op2(e24, e24)) | (e20 = op2(e23, e24)) | (e20 = op2(e22, e24)) | (e20 = op2(e21, e24)) | (e20 = op2(e20, e24))) & ((e20 = op2(e24, e24)) | (e20 = op2(e24, e23)) | (e20 = op2(e24, e22)) | (e20 = op2(e24, e21)) | (e20 = op2(e24, e20))) & ((e24 = op2(e24, e23)) | (e24 = op2(e23, e23)) | (e24 = op2(e22, e23)) | (e24 = op2(e21, e23)) | (e24 = op2(e20, e23))) & ((e24 = op2(e23, e24)) | (e24 = op2(e23, e23)) | (e24 = op2(e23, e22)) | (e24 = op2(e23, e21)) | (e24 = op2(e23, e20))) & ((e23 = op2(e24, e23)) | (e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e24)) | (e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e24, e23)) | (e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e24)) | (e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e24, e23)) | (e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e24)) | (e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e24, e23)) | (e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e24)) | (e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e24 = op2(e24, e22)) | (e24 = op2(e23, e22)) | (e24 = op2(e22, e22)) | (e24 = op2(e21, e22)) | (e24 = op2(e20, e22))) & ((e24 = op2(e22, e24)) | (e24 = op2(e22, e23)) | (e24 = op2(e22, e22)) | (e24 = op2(e22, e21)) | (e24 = op2(e22, e20))) & ((e23 = op2(e24, e22)) | (e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e24)) | (e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e24, e22)) | (e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e24)) | (e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e24, e22)) | (e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e24)) | (e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e24, e22)) | (e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e24)) | (e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e24 = op2(e24, e21)) | (e24 = op2(e23, e21)) | (e24 = op2(e22, e21)) | (e24 = op2(e21, e21)) | (e24 = op2(e20, e21))) & ((e24 = op2(e21, e24)) | (e24 = op2(e21, e23)) | (e24 = op2(e21, e22)) | (e24 = op2(e21, e21)) | (e24 = op2(e21, e20))) & ((e23 = op2(e24, e21)) | (e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e24)) | (e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e24, e21)) | (e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e24)) | (e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e24, e21)) | (e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e24)) | (e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e24, e21)) | (e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e24)) | (e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e24 = op2(e24, e20)) | (e24 = op2(e23, e20)) | (e24 = op2(e22, e20)) | (e24 = op2(e21, e20)) | (op2(e20, e20) = e24)) & ((e24 = op2(e20, e24)) | (e24 = op2(e20, e23)) | (e24 = op2(e20, e22)) | (e24 = op2(e20, e21)) | (op2(e20, e20) = e24)) & ((e23 = op2(e24, e20)) | (e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e24)) | (e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e24, e20)) | (e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e24)) | (e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e24, e20)) | (e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e24)) | (e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e24, e20)) | (e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e24)) | (e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax5)).
fof(f2427, plain, (spl80_237 | spl80_212 | spl80_187 | spl80_162 | spl80_137), inference(avatar_split_clause, [], [f348, f1923, f2028, f2133, f2238, f2343])).
fof(f348, plain, ((e21 = op2(e24, e22)) | (e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))), inference(cnf_transformation, [], [f5])).
fof(f2426, plain, (spl80_198 | spl80_193 | spl80_188 | spl80_183 | spl80_178), inference(avatar_split_clause, [], [f349, f2095, f2116, f2137, f2158, f2179])).
fof(f349, plain, ((e22 = op2(e22, e24)) | (e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))), inference(cnf_transformation, [], [f5])).
fof(f2421, plain, (spl80_240 | spl80_215 | spl80_190 | spl80_165 | spl80_140), inference(avatar_split_clause, [], [f354, f1935, f2040, f2145, f2250, f2355])).
fof(f354, plain, ((e24 = op2(e24, e22)) | (e24 = op2(e23, e22)) | (e24 = op2(e22, e22)) | (e24 = op2(e21, e22)) | (e24 = op2(e20, e22))), inference(cnf_transformation, [], [f5])).
fof(f2402, plain, (spl80_150 | spl80_145 | spl80_140 | spl80_135 | spl80_130), inference(avatar_split_clause, [], [f373, f1893, f1914, f1935, f1956, f1977])).
fof(f373, plain, ((e24 = op2(e24, e24)) | (e24 = op2(e24, e23)) | (e24 = op2(e24, e22)) | (e24 = op2(e24, e21)) | (e24 = op2(e24, e20))), inference(cnf_transformation, [], [f5])).
fof(f2400, plain, (spl80_246 | spl80_247 | spl80_248 | spl80_249 | spl80_250), inference(avatar_split_clause, [], [f300, f2397, f2393, f2389, f2385, f2381])).
fof(f300, plain, ((op2(e20, e20) = e24) | (op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e24 = op2(e24, e24)) | (e23 = op2(e24, e24)) | (e22 = op2(e24, e24)) | (e21 = op2(e24, e24)) | (e20 = op2(e24, e24))) & ((e24 = op2(e24, e23)) | (e23 = op2(e24, e23)) | (e22 = op2(e24, e23)) | (e21 = op2(e24, e23)) | (e20 = op2(e24, e23))) & ((e24 = op2(e24, e22)) | (e23 = op2(e24, e22)) | (e22 = op2(e24, e22)) | (e21 = op2(e24, e22)) | (e20 = op2(e24, e22))) & ((e24 = op2(e24, e21)) | (e23 = op2(e24, e21)) | (e22 = op2(e24, e21)) | (e21 = op2(e24, e21)) | (e20 = op2(e24, e21))) & ((e24 = op2(e24, e20)) | (e23 = op2(e24, e20)) | (e22 = op2(e24, e20)) | (e21 = op2(e24, e20)) | (e20 = op2(e24, e20))) & ((e24 = op2(e23, e24)) | (e23 = op2(e23, e24)) | (e22 = op2(e23, e24)) | (e21 = op2(e23, e24)) | (e20 = op2(e23, e24))) & ((e24 = op2(e23, e23)) | (e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e24 = op2(e23, e22)) | (e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e24 = op2(e23, e21)) | (e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e24 = op2(e23, e20)) | (e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e24 = op2(e22, e24)) | (e23 = op2(e22, e24)) | (e22 = op2(e22, e24)) | (e21 = op2(e22, e24)) | (e20 = op2(e22, e24))) & ((e24 = op2(e22, e23)) | (e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e24 = op2(e22, e22)) | (e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e24 = op2(e22, e21)) | (e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e24 = op2(e22, e20)) | (e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e24 = op2(e21, e24)) | (e23 = op2(e21, e24)) | (e22 = op2(e21, e24)) | (e21 = op2(e21, e24)) | (e20 = op2(e21, e24))) & ((e24 = op2(e21, e23)) | (e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e24 = op2(e21, e22)) | (e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e24 = op2(e21, e21)) | (e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e24 = op2(e21, e20)) | (e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e24 = op2(e20, e24)) | (e23 = op2(e20, e24)) | (e22 = op2(e20, e24)) | (e21 = op2(e20, e24)) | (e20 = op2(e20, e24))) & ((e24 = op2(e20, e23)) | (e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e24 = op2(e20, e22)) | (e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e24 = op2(e20, e21)) | (e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e24) | (op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax4)).
fof(f2274, plain, (spl80_216 | spl80_217 | spl80_218 | spl80_219 | spl80_220), inference(avatar_split_clause, [], [f306, f2271, f2267, f2263, f2259, f2255])).
fof(f306, plain, ((e24 = op2(e21, e21)) | (e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))), inference(cnf_transformation, [], [f4])).
fof(f2022, plain, (spl80_156 | spl80_157 | spl80_158 | spl80_159 | spl80_160), inference(avatar_split_clause, [], [f318, f2019, f2015, f2011, f2007, f2003])).
fof(f318, plain, ((e24 = op2(e23, e23)) | (e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))), inference(cnf_transformation, [], [f4])).
fof(f1867, plain, (spl80_125 | spl80_120 | spl80_115 | spl80_110 | spl80_105), inference(avatar_split_clause, [], [f233, f1738, f1759, f1780, f1801, f1822])).
fof(f233, plain, ((e14 = op1(e10, e14)) | (e14 = op1(e10, e13)) | (e14 = op1(e10, e12)) | (e14 = op1(e10, e11)) | (op1(e10, e10) = e14)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e14 = op1(e14, e14)) | (e14 = op1(e13, e14)) | (e14 = op1(e12, e14)) | (e14 = op1(e11, e14)) | (e14 = op1(e10, e14))) & ((e14 = op1(e14, e14)) | (e14 = op1(e14, e13)) | (e14 = op1(e14, e12)) | (e14 = op1(e14, e11)) | (e14 = op1(e14, e10))) & ((e13 = op1(e14, e14)) | (e13 = op1(e13, e14)) | (e13 = op1(e12, e14)) | (e13 = op1(e11, e14)) | (e13 = op1(e10, e14))) & ((e13 = op1(e14, e14)) | (e13 = op1(e14, e13)) | (e13 = op1(e14, e12)) | (e13 = op1(e14, e11)) | (e13 = op1(e14, e10))) & ((e12 = op1(e14, e14)) | (e12 = op1(e13, e14)) | (e12 = op1(e12, e14)) | (e12 = op1(e11, e14)) | (e12 = op1(e10, e14))) & ((e12 = op1(e14, e14)) | (e12 = op1(e14, e13)) | (e12 = op1(e14, e12)) | (e12 = op1(e14, e11)) | (e12 = op1(e14, e10))) & ((e11 = op1(e14, e14)) | (e11 = op1(e13, e14)) | (e11 = op1(e12, e14)) | (e11 = op1(e11, e14)) | (e11 = op1(e10, e14))) & ((e11 = op1(e14, e14)) | (e11 = op1(e14, e13)) | (e11 = op1(e14, e12)) | (e11 = op1(e14, e11)) | (e11 = op1(e14, e10))) & ((e10 = op1(e14, e14)) | (e10 = op1(e13, e14)) | (e10 = op1(e12, e14)) | (e10 = op1(e11, e14)) | (e10 = op1(e10, e14))) & ((e10 = op1(e14, e14)) | (e10 = op1(e14, e13)) | (e10 = op1(e14, e12)) | (e10 = op1(e14, e11)) | (e10 = op1(e14, e10))) & ((e14 = op1(e14, e13)) | (e14 = op1(e13, e13)) | (e14 = op1(e12, e13)) | (e14 = op1(e11, e13)) | (e14 = op1(e10, e13))) & ((e14 = op1(e13, e14)) | (e14 = op1(e13, e13)) | (e14 = op1(e13, e12)) | (e14 = op1(e13, e11)) | (e14 = op1(e13, e10))) & ((e13 = op1(e14, e13)) | (e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e14)) | (e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e14, e13)) | (e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e14)) | (e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e14, e13)) | (e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e14)) | (e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e14, e13)) | (e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e14)) | (e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e14 = op1(e14, e12)) | (e14 = op1(e13, e12)) | (e14 = op1(e12, e12)) | (e14 = op1(e11, e12)) | (e14 = op1(e10, e12))) & ((e14 = op1(e12, e14)) | (e14 = op1(e12, e13)) | (e14 = op1(e12, e12)) | (e14 = op1(e12, e11)) | (e14 = op1(e12, e10))) & ((e13 = op1(e14, e12)) | (e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e14)) | (e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e14, e12)) | (e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e14)) | (e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e14, e12)) | (e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e14)) | (e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e14, e12)) | (e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e14)) | (e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e14 = op1(e14, e11)) | (e14 = op1(e13, e11)) | (e14 = op1(e12, e11)) | (e14 = op1(e11, e11)) | (e14 = op1(e10, e11))) & ((e14 = op1(e11, e14)) | (e14 = op1(e11, e13)) | (e14 = op1(e11, e12)) | (e14 = op1(e11, e11)) | (e14 = op1(e11, e10))) & ((e13 = op1(e14, e11)) | (e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e14)) | (e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e14, e11)) | (e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e14)) | (e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e14, e11)) | (e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e14)) | (e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e14, e11)) | (e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e14)) | (e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e14 = op1(e14, e10)) | (e14 = op1(e13, e10)) | (e14 = op1(e12, e10)) | (e14 = op1(e11, e10)) | (op1(e10, e10) = e14)) & ((e14 = op1(e10, e14)) | (e14 = op1(e10, e13)) | (e14 = op1(e10, e12)) | (e14 = op1(e10, e11)) | (op1(e10, e10) = e14)) & ((e13 = op1(e14, e10)) | (e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e14)) | (e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e14, e10)) | (e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e14)) | (e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e14, e10)) | (e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e14)) | (e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e14, e10)) | (e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e14)) | (e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax2)).
fof(f1851, plain, (spl80_73 | spl80_68 | spl80_63 | spl80_58 | spl80_53), inference(avatar_split_clause, [], [f249, f1520, f1541, f1562, f1583, f1604])).
fof(f249, plain, ((e12 = op1(e12, e14)) | (e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1847, plain, (spl80_75 | spl80_70 | spl80_65 | spl80_60 | spl80_55), inference(avatar_split_clause, [], [f253, f1528, f1549, f1570, f1591, f1612])).
fof(f253, plain, ((e14 = op1(e12, e14)) | (e14 = op1(e12, e13)) | (e14 = op1(e12, e12)) | (e14 = op1(e12, e11)) | (e14 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1838, plain, (spl80_109 | spl80_84 | spl80_59 | spl80_34 | spl80_9), inference(avatar_split_clause, [], [f262, f1335, f1440, f1545, f1650, f1755])).
fof(f262, plain, ((e13 = op1(e14, e13)) | (e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1834, plain, (spl80_101 | spl80_76 | spl80_51 | spl80_26 | spl80_1), inference(avatar_split_clause, [], [f266, f1302, f1407, f1512, f1617, f1722])).
fof(f266, plain, ((e10 = op1(e14, e14)) | (e10 = op1(e13, e14)) | (e10 = op1(e12, e14)) | (e10 = op1(e11, e14)) | (e10 = op1(e10, e14))), inference(cnf_transformation, [], [f2])).
fof(f1827, plain, (spl80_25 | spl80_20 | spl80_15 | spl80_10 | spl80_5), inference(avatar_split_clause, [], [f273, f1318, f1339, f1360, f1381, f1402])).
fof(f273, plain, ((e14 = op1(e14, e14)) | (e14 = op1(e14, e13)) | (e14 = op1(e14, e12)) | (e14 = op1(e14, e11)) | (e14 = op1(e14, e10))), inference(cnf_transformation, [], [f2])).
fof(f1825, plain, (spl80_121 | spl80_122 | spl80_123 | spl80_124 | spl80_125), inference(avatar_split_clause, [], [f200, f1822, f1818, f1814, f1810, f1806])).
fof(f200, plain, ((op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e14 = op1(e14, e14)) | (e13 = op1(e14, e14)) | (e12 = op1(e14, e14)) | (e11 = op1(e14, e14)) | (e10 = op1(e14, e14))) & ((e14 = op1(e14, e13)) | (e13 = op1(e14, e13)) | (e12 = op1(e14, e13)) | (e11 = op1(e14, e13)) | (e10 = op1(e14, e13))) & ((e14 = op1(e14, e12)) | (e13 = op1(e14, e12)) | (e12 = op1(e14, e12)) | (e11 = op1(e14, e12)) | (e10 = op1(e14, e12))) & ((e14 = op1(e14, e11)) | (e13 = op1(e14, e11)) | (e12 = op1(e14, e11)) | (e11 = op1(e14, e11)) | (e10 = op1(e14, e11))) & ((e14 = op1(e14, e10)) | (e13 = op1(e14, e10)) | (e12 = op1(e14, e10)) | (e11 = op1(e14, e10)) | (e10 = op1(e14, e10))) & ((e14 = op1(e13, e14)) | (e13 = op1(e13, e14)) | (e12 = op1(e13, e14)) | (e11 = op1(e13, e14)) | (e10 = op1(e13, e14))) & ((e14 = op1(e13, e13)) | (e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e14 = op1(e13, e12)) | (e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e14 = op1(e13, e11)) | (e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e14 = op1(e13, e10)) | (e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e14 = op1(e12, e14)) | (e13 = op1(e12, e14)) | (e12 = op1(e12, e14)) | (e11 = op1(e12, e14)) | (e10 = op1(e12, e14))) & ((e14 = op1(e12, e13)) | (e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e14 = op1(e12, e12)) | (e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e14 = op1(e12, e11)) | (e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e14 = op1(e12, e10)) | (e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e14 = op1(e11, e14)) | (e13 = op1(e11, e14)) | (e12 = op1(e11, e14)) | (e11 = op1(e11, e14)) | (e10 = op1(e11, e14))) & ((e14 = op1(e11, e13)) | (e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e14 = op1(e11, e12)) | (e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e14 = op1(e11, e11)) | (e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e14 = op1(e11, e10)) | (e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))) & ((e14 = op1(e10, e13)) | (e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e14 = op1(e10, e12)) | (e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e14 = op1(e10, e11)) | (e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e14) | (op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG168+1.p', ax1)).
fof(f1741, plain, (spl80_101 | spl80_102 | spl80_103 | spl80_104 | spl80_105), inference(avatar_split_clause, [], [f204, f1738, f1734, f1730, f1726, f1722])).
fof(f204, plain, ((e14 = op1(e10, e14)) | (e13 = op1(e10, e14)) | (e12 = op1(e10, e14)) | (e11 = op1(e10, e14)) | (e10 = op1(e10, e14))), inference(cnf_transformation, [], [f1])).
fof(f1699, plain, (spl80_91 | spl80_92 | spl80_93 | spl80_94 | spl80_95), inference(avatar_split_clause, [], [f206, f1696, f1692, f1688, f1684, f1680])).
fof(f206, plain, ((e14 = op1(e11, e11)) | (e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))), inference(cnf_transformation, [], [f1])).