fof(f6334, plain, $false, inference(avatar_sat_refutation, [], [f874, f891, f925, f942, f993, f1010, f1027, f1061, f1078, f1095, f1112, f1129, f1130, f1131, f1132, f1134, f1135, f1136, f1138, f1140, f1141, f1144, f1147, f1148, f1149, f1154, f1155, f1157, f1158, f1160, f1161, f1178, f1195, f1212, f1229, f1297, f1314, f1331, f1365, f1382, f1399, f1416, f1433, f1434, f1435, f1438, f1439, f1440, f1441, f1442, f1444, f1445, f1448, f1449, f1451, f1452, f1455, f1458, f1459, f1462, f1464, f1465, f1474, f1485, f1491, f1493, f1494, f1495, f1500, f1504, f1505, f1515, f1525, f1533, f1534, f1535, f1540, f1541, f1544, f1545, f1555, f1562, f1563, f1565, f1573, f1574, f1575, f1580, f1584, f1585, f1594, f1595, f1605, f1613, f1615, f1620, f1623, f1629, f1636, f1638, f1639, f1647, f1655, f1661, f1671, f1679, f1687, f1693, f1700, f1702, f1710, f1711, f1719, f1725, f1734, f1735, f1736, f1737, f1741, f1744, f1745, f1754, f1765, f1773, f1775, f1780, f1784, f1785, f1795, f1805, f1813, f1814, f1815, f1820, f1821, f1824, f1825, f1835, f1845, f1853, f1854, f1855, f1860, f1864, f1865, f1874, f1875, f1885, f1893, f1895, f1900, f1903, f1909, f1916, f1918, f1927, f1935, f1941, f1951, f1959, f1967, f1975, f1980, f1982, f1990, f1991, f1999, f2005, f2015, f2016, f2017, f2021, f2024, f2025, f2026, f2027, f2071, f2082, f2091, f2097, f2117, f2122, f2135, f2149, f2169, f2180, f2205, f2219, f2225, f2250, f2261, f2362, f2485, f2600, f2710, f2726, f2746, f2765, f2797, f2803, f2827, f2835, f2841, f2853, f2854, f2874, f2877, f2881, f2883, f2888, f2900, f2903, f2904, f2910, f2931, f2943, f2956, f2966, f2997, f2998, f3009, f3016, f3033, f3038, f3060, f3070, f3077, f3081, f3093, f3102, f3110, f3116, f3143, f3164, f3165, f3166, f3168, f3169, f3173, f3177, f3181, f3186, f3207, f3211, f3214, f3239, f3245, f3260, f3265, f3277, f3279, f3306, f3312, f3314, f3318, f3324, f3328, f3330, f3336, f3346, f3347, f3348, f3354, f3364, f3369, f3381, f3400, f3426, f3435, f3438, f3448, f3468, f3470, f3472, f3476, f3496, f3517, f3520, f3537, f3548, f3553, f3569, f3575, f3580, f3603, f3619, f3645, f3681, f3685, f3688, f3705, f3711, f3719, f3744, f3770, f3781, f3792, f3803, f3821, f3837, f3850, f3858, f3862, f3870, f3874, f3883, f3888, f3898, f3901, f3908, f3913, f3921, f3934, f3952, f3956, f3961, f3989, f3995, f4006, f4026, f4029, f4047, f4074, f4077, f4084, f4095, f4110, f4115, f4167, f4195, f4200, f4205, f4238, f4253, f4262, f4272, f4282, f4292, f4301, f4310, f4321, f4324, f4330, f4335, f4340, f4351, f4352, f4357, f4363, f4369, f4416, f4417, f4430, f4435, f4441, f4442, f4445, f4450, f4495, f4508, f4522, f4549, f4555, f4564, f4576, f4585, f4590, f4603, f4609, f4618, f4636, f4643, f4646, f4657, f4672, f4674, f4679, f4681, f4697, f4730, f4763, f4772, f4797, f4798, f4812, f4813, f4830, f4886, f4893, f4895, f4914, f4924, f4930, f4940, f4950, f4959, f4968, f4976, f4982, f4988, f4993, f4998, f5004, f5012, f5058, f5089, f5098, f5117, f5118, f5133, f5137, f5168, f5181, f5182, f5184, f5187, f5194, f5296, f5299, f5352, f5374, f5375, f5376, f5377, f5383, f5388, f5425, f5429, f5438, f5452, f5477, f5483, f5493, f5502, f5514, f5519, f5531, f5537, f5547, f5591, f5592, f5603, f5614, f5617, f5625, f5629, f5639, f5640, f5645, f5676, f5739, f5741, f5759, f5830, f5831, f5840, f5846, f5867, f5879, f5891, f5905, f5918, f5939, f5948, f5953, f5965, f5968, f5978, f5981, f5991, f5996, f6009, f6014, f6026, f6104, f6105, f6107, f6110, f6125, f6136, f6139, f6174, f6175, f6196, f6197, f6213, f6214, f6239, f6257, f6262, f6274, f6280, f6329])).
fof(f6329, plain, (~ spl72_16 | ~ spl72_95 | ~ spl72_190 | ~ spl72_200 | spl72_234), inference(avatar_contradiction_clause, [], [f6328])).
fof(f6328, plain, ($false | (~ spl72_16 | ~ spl72_95 | ~ spl72_190 | ~ spl72_200 | spl72_234)), inference(subsumption_resolution, [], [f6327, f1292])).
fof(f1292, plain, ((e22 = op2(e22, e20)) | ~ spl72_95), inference(avatar_component_clause, [], [f1290])).
fof(f1290, plain, (spl72_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_95])])).
fof(f6327, plain, (~ (e22 = op2(e22, e20)) | (~ spl72_16 | ~ spl72_190 | ~ spl72_200 | spl72_234)), inference(forward_demodulation, [], [f6326, f2069])).
fof(f2069, plain, ((e22 = h4(e13)) | ~ spl72_190), inference(avatar_component_clause, [], [f2068])).
fof(f2068, plain, (spl72_190 <=> (e22 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_190])])).
fof(f6326, plain, (~ (op2(e22, e20) = h4(e13)) | (~ spl72_16 | ~ spl72_190 | ~ spl72_200 | spl72_234)), inference(forward_demodulation, [], [f6325, f924])).
fof(f924, plain, ((e13 = op1(e13, e10)) | ~ spl72_16), inference(avatar_component_clause, [], [f922])).
fof(f922, plain, (spl72_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_16])])).
fof(f6325, plain, (~ (op2(e22, e20) = h4(op1(e13, e10))) | (~ spl72_190 | ~ spl72_200 | spl72_234)), inference(forward_demodulation, [], [f6324, f2069])).
fof(f6324, plain, (~ (h4(op1(e13, e10)) = op2(h4(e13), e20)) | (~ spl72_200 | spl72_234)), inference(forward_demodulation, [], [f2305, f2120])).
fof(f2120, plain, ((e20 = h4(e10)) | ~ spl72_200), inference(avatar_component_clause, [], [f2119])).
fof(f2119, plain, (spl72_200 <=> (e20 = h4(e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_200])])).
fof(f2305, plain, (~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | spl72_234), inference(avatar_component_clause, [], [f2303])).
fof(f2303, plain, (spl72_234 <=> (h4(op1(e13, e10)) = op2(h4(e13), h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_234])])).
fof(f6280, plain, (~ spl72_49 | ~ spl72_117 | ~ spl72_190 | ~ spl72_200 | spl72_231), inference(avatar_contradiction_clause, [], [f6279])).
fof(f6279, plain, ($false | (~ spl72_49 | ~ spl72_117 | ~ spl72_190 | ~ spl72_200 | spl72_231)), inference(subsumption_resolution, [], [f6278, f1386])).
fof(f1386, plain, ((e20 = op2(e20, e22)) | ~ spl72_117), inference(avatar_component_clause, [], [f1384])).
fof(f1384, plain, (spl72_117 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_117])])).
fof(f6278, plain, (~ (e20 = op2(e20, e22)) | (~ spl72_49 | ~ spl72_190 | ~ spl72_200 | spl72_231)), inference(forward_demodulation, [], [f6277, f2120])).
fof(f6277, plain, (~ (op2(e20, e22) = h4(e10)) | (~ spl72_49 | ~ spl72_190 | ~ spl72_200 | spl72_231)), inference(forward_demodulation, [], [f6276, f1065])).
fof(f1065, plain, ((e10 = op1(e10, e13)) | ~ spl72_49), inference(avatar_component_clause, [], [f1063])).
fof(f1063, plain, (spl72_49 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_49])])).
fof(f6276, plain, (~ (op2(e20, e22) = h4(op1(e10, e13))) | (~ spl72_190 | ~ spl72_200 | spl72_231)), inference(forward_demodulation, [], [f6275, f2120])).
fof(f6275, plain, (~ (h4(op1(e10, e13)) = op2(h4(e10), e22)) | (~ spl72_190 | spl72_231)), inference(forward_demodulation, [], [f2293, f2069])).
fof(f2293, plain, (~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | spl72_231), inference(avatar_component_clause, [], [f2291])).
fof(f2291, plain, (spl72_231 <=> (h4(op1(e10, e13)) = op2(h4(e10), h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_231])])).
fof(f6274, plain, (~ spl72_19 | ~ spl72_72 | ~ spl72_190 | spl72_239), inference(avatar_contradiction_clause, [], [f6273])).
fof(f6273, plain, ($false | (~ spl72_19 | ~ spl72_72 | ~ spl72_190 | spl72_239)), inference(subsumption_resolution, [], [f6272, f1194])).
fof(f1194, plain, ((e23 = op2(e23, e22)) | ~ spl72_72), inference(avatar_component_clause, [], [f1192])).
fof(f1192, plain, (spl72_72 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_72])])).
fof(f6272, plain, (~ (e23 = op2(e23, e22)) | (~ spl72_19 | ~ spl72_190 | spl72_239)), inference(forward_demodulation, [], [f6271, f758])).
fof(f758, plain, (e23 = h4(e12)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((h4(e13) = op2(e23, op2(e23, e23))) & (op2(e23, e23) = h4(e11)) & (h4(e10) = op2(op2(e23, e23), op2(e23, e23))) & (e23 = h4(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax17)).
fof(f6271, plain, (~ (op2(e23, e22) = h4(e12)) | (~ spl72_19 | ~ spl72_190 | spl72_239)), inference(forward_demodulation, [], [f6270, f937])).
fof(f937, plain, ((e12 = op1(e12, e13)) | ~ spl72_19), inference(avatar_component_clause, [], [f935])).
fof(f935, plain, (spl72_19 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_19])])).
fof(f6270, plain, (~ (op2(e23, e22) = h4(op1(e12, e13))) | (~ spl72_190 | spl72_239)), inference(forward_demodulation, [], [f2325, f2069])).
fof(f2325, plain, (~ (h4(op1(e12, e13)) = op2(e23, h4(e13))) | spl72_239), inference(avatar_component_clause, [], [f2323])).
fof(f2323, plain, (spl72_239 <=> (h4(op1(e12, e13)) = op2(e23, h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_239])])).
fof(f6262, plain, (~ spl72_5 | ~ spl72_81 | ~ spl72_190 | ~ spl72_200 | spl72_238), inference(avatar_contradiction_clause, [], [f6261])).
fof(f6261, plain, ($false | (~ spl72_5 | ~ spl72_81 | ~ spl72_190 | ~ spl72_200 | spl72_238)), inference(subsumption_resolution, [], [f6260, f1233])).
fof(f1233, plain, ((e20 = op2(e22, e23)) | ~ spl72_81), inference(avatar_component_clause, [], [f1231])).
fof(f1231, plain, (spl72_81 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_81])])).
fof(f6260, plain, (~ (e20 = op2(e22, e23)) | (~ spl72_5 | ~ spl72_190 | ~ spl72_200 | spl72_238)), inference(forward_demodulation, [], [f6259, f2120])).
fof(f6259, plain, (~ (op2(e22, e23) = h4(e10)) | (~ spl72_5 | ~ spl72_190 | spl72_238)), inference(forward_demodulation, [], [f6258, f878])).
fof(f878, plain, ((e10 = op1(e13, e12)) | ~ spl72_5), inference(avatar_component_clause, [], [f876])).
fof(f876, plain, (spl72_5 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_5])])).
fof(f6258, plain, (~ (op2(e22, e23) = h4(op1(e13, e12))) | (~ spl72_190 | spl72_238)), inference(forward_demodulation, [], [f2321, f2069])).
fof(f2321, plain, (~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | spl72_238), inference(avatar_component_clause, [], [f2319])).
fof(f2319, plain, (spl72_238 <=> (h4(op1(e13, e12)) = op2(h4(e13), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_238])])).
fof(f6257, plain, (~ spl72_56 | ~ spl72_115 | ~ spl72_190 | ~ spl72_200 | spl72_245), inference(avatar_contradiction_clause, [], [f6256])).
fof(f6256, plain, ($false | (~ spl72_56 | ~ spl72_115 | ~ spl72_190 | ~ spl72_200 | spl72_245)), inference(subsumption_resolution, [], [f6255, f1377])).
fof(f1377, plain, ((e22 = op2(e20, e23)) | ~ spl72_115), inference(avatar_component_clause, [], [f1375])).
fof(f1375, plain, (spl72_115 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_115])])).
fof(f6255, plain, (~ (e22 = op2(e20, e23)) | (~ spl72_56 | ~ spl72_190 | ~ spl72_200 | spl72_245)), inference(forward_demodulation, [], [f6254, f2069])).
fof(f6254, plain, (~ (op2(e20, e23) = h4(e13)) | (~ spl72_56 | ~ spl72_200 | spl72_245)), inference(forward_demodulation, [], [f6029, f2120])).
fof(f6029, plain, (~ (h4(e13) = op2(h4(e10), e23)) | (~ spl72_56 | spl72_245)), inference(forward_demodulation, [], [f2349, f1094])).
fof(f1094, plain, ((e13 = op1(e10, e12)) | ~ spl72_56), inference(avatar_component_clause, [], [f1092])).
fof(f1092, plain, (spl72_56 <=> (e13 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_56])])).
fof(f2349, plain, (~ (h4(op1(e10, e12)) = op2(h4(e10), e23)) | spl72_245), inference(avatar_component_clause, [], [f2347])).
fof(f2347, plain, (spl72_245 <=> (h4(op1(e10, e12)) = op2(h4(e10), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_245])])).
fof(f6239, plain, (~ spl72_83 | ~ spl72_95), inference(avatar_split_clause, [], [f6238, f1290, f1239])).
fof(f1239, plain, (spl72_83 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_83])])).
fof(f6238, plain, (~ (e22 = op2(e22, e23)) | ~ spl72_95), inference(backward_demodulation, [], [f350, f1292])).
fof(f350, plain, ~ (op2(e22, e20) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax6)).
fof(f6214, plain, (~ spl72_78 | ~ spl72_195), inference(avatar_split_clause, [], [f6202, f2094, f1218])).
fof(f1218, plain, (spl72_78 <=> (e21 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_78])])).
fof(f2094, plain, (spl72_195 <=> (e21 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_195])])).
fof(f6202, plain, (~ (e21 = op2(e23, e20)) | ~ spl72_195), inference(backward_demodulation, [], [f2058, f2095])).
fof(f2095, plain, ((e21 = h4(e11)) | ~ spl72_195), inference(avatar_component_clause, [], [f2094])).
fof(f2058, plain, ~ (op2(e23, e20) = h4(e11)), inference(backward_demodulation, [], [f356, f760])).
fof(f760, plain, (op2(e23, e23) = h4(e11)), inference(cnf_transformation, [], [f17])).
fof(f356, plain, ~ (op2(e23, e20) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f6213, plain, (~ spl72_98 | ~ spl72_195), inference(avatar_split_clause, [], [f6200, f2094, f1303])).
fof(f1303, plain, (spl72_98 <=> (e21 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_98])])).
fof(f6200, plain, (~ (e21 = op2(e21, e23)) | ~ spl72_195), inference(backward_demodulation, [], [f2056, f2095])).
fof(f2056, plain, ~ (op2(e21, e23) = h4(e11)), inference(backward_demodulation, [], [f334, f760])).
fof(f334, plain, ~ (op2(e21, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f6197, plain, (spl72_128 | ~ spl72_63 | ~ spl72_200 | ~ spl72_229), inference(avatar_split_clause, [], [f6195, f2283, f2119, f1122, f1430])).
fof(f1430, plain, (spl72_128 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl72_128])])).
fof(f1122, plain, (spl72_63 <=> (op1(e10, e10) = e12)), introduced(avatar_definition, [new_symbols(naming, [spl72_63])])).
fof(f2283, plain, (spl72_229 <=> (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_229])])).
fof(f6195, plain, ((op2(e20, e20) = e23) | (~ spl72_63 | ~ spl72_200 | ~ spl72_229)), inference(backward_demodulation, [], [f6146, f2120])).
fof(f6146, plain, ((e23 = op2(h4(e10), h4(e10))) | (~ spl72_63 | ~ spl72_229)), inference(forward_demodulation, [], [f6145, f758])).
fof(f6145, plain, ((h4(e12) = op2(h4(e10), h4(e10))) | (~ spl72_63 | ~ spl72_229)), inference(forward_demodulation, [], [f2284, f1124])).
fof(f1124, plain, ((op1(e10, e10) = e12) | ~ spl72_63), inference(avatar_component_clause, [], [f1122])).
fof(f2284, plain, ((h4(op1(e10, e10)) = op2(h4(e10), h4(e10))) | ~ spl72_229), inference(avatar_component_clause, [], [f2283])).
fof(f6196, plain, (~ spl72_77 | ~ spl72_29 | ~ spl72_200 | spl72_242), inference(avatar_split_clause, [], [f6193, f2335, f2119, f978, f1214])).
fof(f1214, plain, (spl72_77 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_77])])).
fof(f978, plain, (spl72_29 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_29])])).
fof(f2335, plain, (spl72_242 <=> (h4(op1(e12, e10)) = op2(e23, h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_242])])).
fof(f6193, plain, (~ (e20 = op2(e23, e20)) | (~ spl72_29 | ~ spl72_200 | spl72_242)), inference(backward_demodulation, [], [f6033, f2120])).
fof(f6033, plain, (~ (h4(e10) = op2(e23, h4(e10))) | (~ spl72_29 | spl72_242)), inference(forward_demodulation, [], [f2337, f980])).
fof(f980, plain, ((e10 = op1(e12, e10)) | ~ spl72_29), inference(avatar_component_clause, [], [f978])).
fof(f2337, plain, (~ (h4(op1(e12, e10)) = op2(e23, h4(e10))) | spl72_242), inference(avatar_component_clause, [], [f2335])).
fof(f6175, plain, (~ spl72_98 | ~ spl72_210 | ~ spl72_215), inference(avatar_split_clause, [], [f6165, f2202, f2175, f1303])).
fof(f2175, plain, (spl72_210 <=> (e21 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_210])])).
fof(f2202, plain, (spl72_215 <=> (e20 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_215])])).
fof(f6165, plain, (~ (e21 = op2(e21, e23)) | (~ spl72_210 | ~ spl72_215)), inference(backward_demodulation, [], [f3350, f2176])).
fof(f2176, plain, ((e21 = h2(e13)) | ~ spl72_210), inference(avatar_component_clause, [], [f2175])).
fof(f3350, plain, (~ (op2(e21, e23) = h2(e13)) | ~ spl72_215), inference(backward_demodulation, [], [f344, f3341])).
fof(f3341, plain, ((op2(e21, e20) = h2(e13)) | ~ spl72_215), inference(backward_demodulation, [], [f2042, f2203])).
fof(f2203, plain, ((e20 = h2(e11)) | ~ spl72_215), inference(avatar_component_clause, [], [f2202])).
fof(f2042, plain, (h2(e13) = op2(e21, h2(e11))), inference(backward_demodulation, [], [f753, f752])).
fof(f752, plain, (op2(e21, e21) = h2(e11)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((h2(e13) = op2(e21, op2(e21, e21))) & (op2(e21, e21) = h2(e11)) & (h2(e10) = op2(op2(e21, e21), op2(e21, e21))) & (e21 = h2(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax15)).
fof(f753, plain, (h2(e13) = op2(e21, op2(e21, e21))), inference(cnf_transformation, [], [f15])).
fof(f344, plain, ~ (op2(e21, e20) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f6174, plain, (~ spl72_78 | ~ spl72_210 | ~ spl72_215), inference(avatar_split_clause, [], [f6164, f2202, f2175, f1218])).
fof(f6164, plain, (~ (e21 = op2(e23, e20)) | (~ spl72_210 | ~ spl72_215)), inference(backward_demodulation, [], [f3349, f2176])).
fof(f3349, plain, (~ (op2(e23, e20) = h2(e13)) | ~ spl72_215), inference(backward_demodulation, [], [f316, f3341])).
fof(f316, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f6139, plain, (~ spl72_83 | ~ spl72_19 | ~ spl72_39 | ~ spl72_103 | ~ spl72_202 | spl72_252 | ~ spl72_257), inference(avatar_split_clause, [], [f6138, f2458, f2438, f2132, f1324, f1020, f935, f1239])).
fof(f1020, plain, (spl72_39 <=> (e12 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_39])])).
fof(f1324, plain, (spl72_103 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_103])])).
fof(f2132, plain, (spl72_202 <=> (e21 = h3(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_202])])).
fof(f2438, plain, (spl72_252 <=> (op2(e22, e23) = h3(op1(e12, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_252])])).
fof(f2458, plain, (spl72_257 <=> (h3(op1(e11, e12)) = op2(h3(e11), e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_257])])).
fof(f6138, plain, (~ (e22 = op2(e22, e23)) | (~ spl72_19 | ~ spl72_39 | ~ spl72_103 | ~ spl72_202 | spl72_252 | ~ spl72_257)), inference(forward_demodulation, [], [f6137, f4134])).
fof(f4134, plain, ((e22 = h3(e12)) | (~ spl72_39 | ~ spl72_103 | ~ spl72_202 | ~ spl72_257)), inference(backward_demodulation, [], [f4097, f1022])).
fof(f1022, plain, ((e12 = op1(e11, e12)) | ~ spl72_39), inference(avatar_component_clause, [], [f1020])).
fof(f4097, plain, ((e22 = h3(op1(e11, e12))) | (~ spl72_103 | ~ spl72_202 | ~ spl72_257)), inference(forward_demodulation, [], [f4096, f1326])).
fof(f1326, plain, ((e22 = op2(e21, e22)) | ~ spl72_103), inference(avatar_component_clause, [], [f1324])).
fof(f4096, plain, ((op2(e21, e22) = h3(op1(e11, e12))) | (~ spl72_202 | ~ spl72_257)), inference(forward_demodulation, [], [f2459, f2133])).
fof(f2133, plain, ((e21 = h3(e11)) | ~ spl72_202), inference(avatar_component_clause, [], [f2132])).
fof(f2459, plain, ((h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ spl72_257), inference(avatar_component_clause, [], [f2458])).
fof(f6137, plain, (~ (op2(e22, e23) = h3(e12)) | (~ spl72_19 | spl72_252)), inference(forward_demodulation, [], [f2440, f937])).
fof(f2440, plain, (~ (op2(e22, e23) = h3(op1(e12, e13))) | spl72_252), inference(avatar_component_clause, [], [f2438])).
fof(f6136, plain, (~ spl72_93 | ~ spl72_29 | spl72_255), inference(avatar_split_clause, [], [f6135, f2450, f978, f1282])).
fof(f1282, plain, (spl72_93 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_93])])).
fof(f2450, plain, (spl72_255 <=> (op2(e22, e20) = h3(op1(e12, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_255])])).
fof(f6135, plain, (~ (e20 = op2(e22, e20)) | (~ spl72_29 | spl72_255)), inference(forward_demodulation, [], [f6134, f2054])).
fof(f2054, plain, (e20 = h3(e10)), inference(forward_demodulation, [], [f2053, f2051])).
fof(f2051, plain, (e20 = op2(h3(e11), h3(e11))), inference(backward_demodulation, [], [f743, f756])).
fof(f756, plain, (op2(e22, e22) = h3(e11)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(e22, op2(e22, e22)) = h3(e13)) & (op2(e22, e22) = h3(e11)) & (op2(op2(e22, e22), op2(e22, e22)) = h3(e10)) & (e22 = h3(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax16)).
fof(f743, plain, (e20 = op2(op2(e22, e22), op2(e22, e22))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e23 = op2(e22, op2(e22, e22))) & (e21 = op2(e22, e22)) & (e20 = op2(op2(e22, e22), op2(e22, e22)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax13)).
fof(f2053, plain, (h3(e10) = op2(h3(e11), h3(e11))), inference(forward_demodulation, [], [f755, f756])).
fof(f755, plain, (op2(op2(e22, e22), op2(e22, e22)) = h3(e10)), inference(cnf_transformation, [], [f16])).
fof(f6134, plain, (~ (op2(e22, e20) = h3(e10)) | (~ spl72_29 | spl72_255)), inference(forward_demodulation, [], [f2452, f980])).
fof(f2452, plain, (~ (op2(e22, e20) = h3(op1(e12, e10))) | spl72_255), inference(avatar_component_clause, [], [f2450])).
fof(f6125, plain, (~ spl72_219 | ~ spl72_39 | ~ spl72_63 | ~ spl72_103 | ~ spl72_202 | ~ spl72_257 | spl72_263), inference(avatar_split_clause, [], [f6124, f2482, f2458, f2132, f1324, f1122, f1020, f2222])).
fof(f2222, plain, (spl72_219 <=> (e22 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_219])])).
fof(f2482, plain, (spl72_263 <=> (h1(e11) = h3(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_263])])).
fof(f6124, plain, (~ (e22 = h1(e11)) | (~ spl72_39 | ~ spl72_63 | ~ spl72_103 | ~ spl72_202 | ~ spl72_257 | spl72_263)), inference(forward_demodulation, [], [f6123, f4134])).
fof(f6123, plain, (~ (h1(e11) = h3(e12)) | (~ spl72_63 | spl72_263)), inference(forward_demodulation, [], [f2484, f1124])).
fof(f2484, plain, (~ (h1(e11) = h3(op1(e10, e10))) | spl72_263), inference(avatar_component_clause, [], [f2482])).
fof(f6110, plain, (spl72_77 | ~ spl72_39 | ~ spl72_298 | ~ spl72_301), inference(avatar_split_clause, [], [f6096, f2723, f2699, f1020, f1214])).
fof(f2699, plain, (spl72_298 <=> (h1(op1(e11, e12)) = op2(h1(e11), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_298])])).
fof(f2723, plain, (spl72_301 <=> (e23 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_301])])).
fof(f6096, plain, ((e20 = op2(e23, e20)) | (~ spl72_39 | ~ spl72_298 | ~ spl72_301)), inference(backward_demodulation, [], [f6060, f2724])).
fof(f2724, plain, ((e23 = h1(e11)) | ~ spl72_301), inference(avatar_component_clause, [], [f2723])).
fof(f6060, plain, ((e20 = op2(h1(e11), e20)) | (~ spl72_39 | ~ spl72_298)), inference(forward_demodulation, [], [f6059, f746])).
fof(f746, plain, (e20 = h1(e12)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((h1(e13) = op2(e20, op2(e20, e20))) & (op2(e20, e20) = h1(e11)) & (h1(e10) = op2(op2(e20, e20), op2(e20, e20))) & (e20 = h1(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax14)).
fof(f6059, plain, ((h1(e12) = op2(h1(e11), e20)) | (~ spl72_39 | ~ spl72_298)), inference(forward_demodulation, [], [f2700, f1022])).
fof(f2700, plain, ((h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ spl72_298), inference(avatar_component_clause, [], [f2699])).
fof(f6107, plain, (spl72_66 | ~ spl72_224 | ~ spl72_301), inference(avatar_split_clause, [], [f6093, f2723, f2247, f1167])).
fof(f1167, plain, (spl72_66 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_66])])).
fof(f2247, plain, (spl72_224 <=> (e21 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_224])])).
fof(f6093, plain, ((e21 = op2(e23, e23)) | (~ spl72_224 | ~ spl72_301)), inference(backward_demodulation, [], [f5340, f2724])).
fof(f5340, plain, ((e21 = op2(h1(e11), h1(e11))) | ~ spl72_224), inference(backward_demodulation, [], [f2035, f2248])).
fof(f2248, plain, ((e21 = h1(e10)) | ~ spl72_224), inference(avatar_component_clause, [], [f2247])).
fof(f2035, plain, (h1(e10) = op2(h1(e11), h1(e11))), inference(forward_demodulation, [], [f747, f748])).
fof(f748, plain, (op2(e20, e20) = h1(e11)), inference(cnf_transformation, [], [f14])).
fof(f747, plain, (h1(e10) = op2(op2(e20, e20), op2(e20, e20))), inference(cnf_transformation, [], [f14])).
fof(f6105, plain, (~ spl72_116 | ~ spl72_301), inference(avatar_split_clause, [], [f6091, f2723, f1379])).
fof(f1379, plain, (spl72_116 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_116])])).
fof(f6091, plain, (~ (e23 = op2(e20, e23)) | ~ spl72_301), inference(backward_demodulation, [], [f2033, f2724])).
fof(f2033, plain, ~ (op2(e20, e23) = h1(e11)), inference(backward_demodulation, [], [f338, f748])).
fof(f338, plain, ~ (op2(e20, e20) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f6104, plain, (~ spl72_112 | ~ spl72_301), inference(avatar_split_clause, [], [f6088, f2723, f1362])).
fof(f1362, plain, (spl72_112 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_112])])).
fof(f6088, plain, (~ (e23 = op2(e21, e20)) | ~ spl72_301), inference(backward_demodulation, [], [f2028, f2724])).
fof(f2028, plain, ~ (op2(e21, e20) = h1(e11)), inference(backward_demodulation, [], [f312, f748])).
fof(f312, plain, ~ (op2(e20, e20) = op2(e21, e20)), inference(cnf_transformation, [], [f6])).
fof(f6026, plain, (~ spl72_2 | ~ spl72_65 | ~ spl72_215 | spl72_272 | ~ spl72_273), inference(avatar_contradiction_clause, [], [f6025])).
fof(f6025, plain, ($false | (~ spl72_2 | ~ spl72_65 | ~ spl72_215 | spl72_272 | ~ spl72_273)), inference(subsumption_resolution, [], [f6024, f1165])).
fof(f1165, plain, ((e20 = op2(e23, e23)) | ~ spl72_65), inference(avatar_component_clause, [], [f1163])).
fof(f1163, plain, (spl72_65 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_65])])).
fof(f6024, plain, (~ (e20 = op2(e23, e23)) | (~ spl72_2 | ~ spl72_215 | spl72_272 | ~ spl72_273)), inference(forward_demodulation, [], [f6023, f2203])).
fof(f6023, plain, (~ (op2(e23, e23) = h2(e11)) | (~ spl72_2 | spl72_272 | ~ spl72_273)), inference(forward_demodulation, [], [f6022, f865])).
fof(f865, plain, ((e11 = op1(e13, e13)) | ~ spl72_2), inference(avatar_component_clause, [], [f863])).
fof(f863, plain, (spl72_2 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_2])])).
fof(f6022, plain, (~ (op2(e23, e23) = h2(op1(e13, e13))) | (spl72_272 | ~ spl72_273)), inference(forward_demodulation, [], [f2563, f2566])).
fof(f2566, plain, ((e23 = h2(e13)) | ~ spl72_273), inference(avatar_component_clause, [], [f2565])).
fof(f2565, plain, (spl72_273 <=> (e23 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_273])])).
fof(f2563, plain, (~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | spl72_272), inference(avatar_component_clause, [], [f2561])).
fof(f2561, plain, (spl72_272 <=> (h2(op1(e13, e13)) = op2(h2(e13), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_272])])).
fof(f6014, plain, (~ spl72_5 | ~ spl72_75 | ~ spl72_208 | ~ spl72_273 | spl72_274), inference(avatar_contradiction_clause, [], [f6013])).
fof(f6013, plain, ($false | (~ spl72_5 | ~ spl72_75 | ~ spl72_208 | ~ spl72_273 | spl72_274)), inference(subsumption_resolution, [], [f6012, f1207])).
fof(f1207, plain, ((e22 = op2(e23, e21)) | ~ spl72_75), inference(avatar_component_clause, [], [f1205])).
fof(f1205, plain, (spl72_75 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_75])])).
fof(f6012, plain, (~ (e22 = op2(e23, e21)) | (~ spl72_5 | ~ spl72_208 | ~ spl72_273 | spl72_274)), inference(forward_demodulation, [], [f6011, f2167])).
fof(f2167, plain, ((e22 = h2(e10)) | ~ spl72_208), inference(avatar_component_clause, [], [f2166])).
fof(f2166, plain, (spl72_208 <=> (e22 = h2(e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_208])])).
fof(f6011, plain, (~ (op2(e23, e21) = h2(e10)) | (~ spl72_5 | ~ spl72_273 | spl72_274)), inference(forward_demodulation, [], [f6010, f878])).
fof(f6010, plain, (~ (op2(e23, e21) = h2(op1(e13, e12))) | (~ spl72_273 | spl72_274)), inference(forward_demodulation, [], [f2571, f2566])).
fof(f2571, plain, (~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | spl72_274), inference(avatar_component_clause, [], [f2569])).
fof(f2569, plain, (spl72_274 <=> (h2(op1(e13, e12)) = op2(h2(e13), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_274])])).
fof(f6009, plain, (~ spl72_16 | ~ spl72_72 | ~ spl72_208 | spl72_270 | ~ spl72_273), inference(avatar_contradiction_clause, [], [f6008])).
fof(f6008, plain, ($false | (~ spl72_16 | ~ spl72_72 | ~ spl72_208 | spl72_270 | ~ spl72_273)), inference(subsumption_resolution, [], [f6007, f1194])).
fof(f6007, plain, (~ (e23 = op2(e23, e22)) | (~ spl72_16 | ~ spl72_208 | spl72_270 | ~ spl72_273)), inference(forward_demodulation, [], [f6006, f2566])).
fof(f6006, plain, (~ (op2(e23, e22) = h2(e13)) | (~ spl72_16 | ~ spl72_208 | spl72_270 | ~ spl72_273)), inference(forward_demodulation, [], [f6005, f924])).
fof(f6005, plain, (~ (op2(e23, e22) = h2(op1(e13, e10))) | (~ spl72_208 | spl72_270 | ~ spl72_273)), inference(forward_demodulation, [], [f6004, f2566])).
fof(f6004, plain, (~ (h2(op1(e13, e10)) = op2(h2(e13), e22)) | (~ spl72_208 | spl72_270)), inference(forward_demodulation, [], [f2555, f2167])).
fof(f2555, plain, (~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | spl72_270), inference(avatar_component_clause, [], [f2553])).
fof(f2553, plain, (spl72_270 <=> (h2(op1(e13, e10)) = op2(h2(e13), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_270])])).
fof(f5996, plain, (~ spl72_19 | ~ spl72_98 | ~ spl72_273 | spl72_275), inference(avatar_contradiction_clause, [], [f5995])).
fof(f5995, plain, ($false | (~ spl72_19 | ~ spl72_98 | ~ spl72_273 | spl72_275)), inference(subsumption_resolution, [], [f5994, f1305])).
fof(f1305, plain, ((e21 = op2(e21, e23)) | ~ spl72_98), inference(avatar_component_clause, [], [f1303])).
fof(f5994, plain, (~ (e21 = op2(e21, e23)) | (~ spl72_19 | ~ spl72_273 | spl72_275)), inference(forward_demodulation, [], [f5993, f750])).
fof(f750, plain, (e21 = h2(e12)), inference(cnf_transformation, [], [f15])).
fof(f5993, plain, (~ (op2(e21, e23) = h2(e12)) | (~ spl72_19 | ~ spl72_273 | spl72_275)), inference(forward_demodulation, [], [f5992, f937])).
fof(f5992, plain, (~ (op2(e21, e23) = h2(op1(e12, e13))) | (~ spl72_273 | spl72_275)), inference(forward_demodulation, [], [f2575, f2566])).
fof(f2575, plain, (~ (h2(op1(e12, e13)) = op2(e21, h2(e13))) | spl72_275), inference(avatar_component_clause, [], [f2573])).
fof(f2573, plain, (spl72_275 <=> (h2(op1(e12, e13)) = op2(e21, h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_275])])).
fof(f5991, plain, (~ spl72_36 | ~ spl72_116 | ~ spl72_215 | spl72_269 | ~ spl72_273), inference(avatar_contradiction_clause, [], [f5990])).
fof(f5990, plain, ($false | (~ spl72_36 | ~ spl72_116 | ~ spl72_215 | spl72_269 | ~ spl72_273)), inference(subsumption_resolution, [], [f5989, f1381])).
fof(f1381, plain, ((e23 = op2(e20, e23)) | ~ spl72_116), inference(avatar_component_clause, [], [f1379])).
fof(f5989, plain, (~ (e23 = op2(e20, e23)) | (~ spl72_36 | ~ spl72_215 | spl72_269 | ~ spl72_273)), inference(forward_demodulation, [], [f5988, f2566])).
fof(f5988, plain, (~ (op2(e20, e23) = h2(e13)) | (~ spl72_36 | ~ spl72_215 | spl72_269 | ~ spl72_273)), inference(forward_demodulation, [], [f5987, f1009])).
fof(f1009, plain, ((e13 = op1(e11, e13)) | ~ spl72_36), inference(avatar_component_clause, [], [f1007])).
fof(f1007, plain, (spl72_36 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_36])])).
fof(f5987, plain, (~ (op2(e20, e23) = h2(op1(e11, e13))) | (~ spl72_215 | spl72_269 | ~ spl72_273)), inference(forward_demodulation, [], [f5986, f2203])).
fof(f5986, plain, (~ (h2(op1(e11, e13)) = op2(h2(e11), e23)) | (spl72_269 | ~ spl72_273)), inference(forward_demodulation, [], [f2551, f2566])).
fof(f2551, plain, (~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | spl72_269), inference(avatar_component_clause, [], [f2549])).
fof(f2549, plain, (spl72_269 <=> (h2(op1(e11, e13)) = op2(h2(e11), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_269])])).
fof(f5981, plain, (~ spl72_22 | spl72_276), inference(avatar_contradiction_clause, [], [f5980])).
fof(f5980, plain, ($false | (~ spl72_22 | spl72_276)), inference(trivial_inequality_removal, [], [f5979])).
fof(f5979, plain, (~ (h2(e11) = h2(e11)) | (~ spl72_22 | spl72_276)), inference(forward_demodulation, [], [f2579, f950])).
fof(f950, plain, ((e11 = op1(e12, e12)) | ~ spl72_22), inference(avatar_component_clause, [], [f948])).
fof(f948, plain, (spl72_22 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_22])])).
fof(f2579, plain, (~ (h2(e11) = h2(op1(e12, e12))) | spl72_276), inference(avatar_component_clause, [], [f2577])).
fof(f2577, plain, (spl72_276 <=> (h2(e11) = h2(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_276])])).
fof(f5978, plain, (~ spl72_46 | ~ spl72_117 | ~ spl72_208 | ~ spl72_215 | spl72_268), inference(avatar_contradiction_clause, [], [f5977])).
fof(f5977, plain, ($false | (~ spl72_46 | ~ spl72_117 | ~ spl72_208 | ~ spl72_215 | spl72_268)), inference(subsumption_resolution, [], [f5976, f1386])).
fof(f5976, plain, (~ (e20 = op2(e20, e22)) | (~ spl72_46 | ~ spl72_208 | ~ spl72_215 | spl72_268)), inference(forward_demodulation, [], [f5975, f2203])).
fof(f5975, plain, (~ (op2(e20, e22) = h2(e11)) | (~ spl72_46 | ~ spl72_208 | ~ spl72_215 | spl72_268)), inference(forward_demodulation, [], [f5974, f1052])).
fof(f1052, plain, ((e11 = op1(e11, e10)) | ~ spl72_46), inference(avatar_component_clause, [], [f1050])).
fof(f1050, plain, (spl72_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_46])])).
fof(f5974, plain, (~ (op2(e20, e22) = h2(op1(e11, e10))) | (~ spl72_208 | ~ spl72_215 | spl72_268)), inference(forward_demodulation, [], [f5973, f2203])).
fof(f5973, plain, (~ (h2(op1(e11, e10)) = op2(h2(e11), e22)) | (~ spl72_208 | spl72_268)), inference(forward_demodulation, [], [f2547, f2167])).
fof(f2547, plain, (~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | spl72_268), inference(avatar_component_clause, [], [f2545])).
fof(f2545, plain, (spl72_268 <=> (h2(op1(e11, e10)) = op2(h2(e11), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_268])])).
fof(f5968, plain, (~ spl72_28 | spl72_277), inference(avatar_contradiction_clause, [], [f5967])).
fof(f5967, plain, ($false | (~ spl72_28 | spl72_277)), inference(trivial_inequality_removal, [], [f5966])).
fof(f5966, plain, (~ (h2(e13) = h2(e13)) | (~ spl72_28 | spl72_277)), inference(forward_demodulation, [], [f2583, f975])).
fof(f975, plain, ((e13 = op1(e12, e11)) | ~ spl72_28), inference(avatar_component_clause, [], [f973])).
fof(f973, plain, (spl72_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_28])])).
fof(f2583, plain, (~ (h2(e13) = h2(op1(e12, e11))) | spl72_277), inference(avatar_component_clause, [], [f2581])).
fof(f2581, plain, (spl72_277 <=> (h2(e13) = h2(op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_277])])).
fof(f5965, plain, (~ spl72_58 | ~ spl72_93 | ~ spl72_208 | ~ spl72_215 | spl72_266), inference(avatar_contradiction_clause, [], [f5964])).
fof(f5964, plain, ($false | (~ spl72_58 | ~ spl72_93 | ~ spl72_208 | ~ spl72_215 | spl72_266)), inference(subsumption_resolution, [], [f5963, f1284])).
fof(f1284, plain, ((e20 = op2(e22, e20)) | ~ spl72_93), inference(avatar_component_clause, [], [f1282])).
fof(f5963, plain, (~ (e20 = op2(e22, e20)) | (~ spl72_58 | ~ spl72_208 | ~ spl72_215 | spl72_266)), inference(forward_demodulation, [], [f5962, f2203])).
fof(f5962, plain, (~ (op2(e22, e20) = h2(e11)) | (~ spl72_58 | ~ spl72_208 | ~ spl72_215 | spl72_266)), inference(forward_demodulation, [], [f5961, f1103])).
fof(f1103, plain, ((e11 = op1(e10, e11)) | ~ spl72_58), inference(avatar_component_clause, [], [f1101])).
fof(f1101, plain, (spl72_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_58])])).
fof(f5961, plain, (~ (op2(e22, e20) = h2(op1(e10, e11))) | (~ spl72_208 | ~ spl72_215 | spl72_266)), inference(forward_demodulation, [], [f5960, f2167])).
fof(f5960, plain, (~ (h2(op1(e10, e11)) = op2(h2(e10), e20)) | (~ spl72_215 | spl72_266)), inference(forward_demodulation, [], [f2539, f2203])).
fof(f2539, plain, (~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | spl72_266), inference(avatar_component_clause, [], [f2537])).
fof(f2537, plain, (spl72_266 <=> (h2(op1(e10, e11)) = op2(h2(e10), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_266])])).
fof(f5953, plain, (~ spl72_39 | ~ spl72_122 | ~ spl72_215 | spl72_279), inference(avatar_contradiction_clause, [], [f5952])).
fof(f5952, plain, ($false | (~ spl72_39 | ~ spl72_122 | ~ spl72_215 | spl72_279)), inference(subsumption_resolution, [], [f5951, f1407])).
fof(f1407, plain, ((e21 = op2(e20, e21)) | ~ spl72_122), inference(avatar_component_clause, [], [f1405])).
fof(f1405, plain, (spl72_122 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_122])])).
fof(f5951, plain, (~ (e21 = op2(e20, e21)) | (~ spl72_39 | ~ spl72_215 | spl72_279)), inference(forward_demodulation, [], [f5950, f750])).
fof(f5950, plain, (~ (op2(e20, e21) = h2(e12)) | (~ spl72_39 | ~ spl72_215 | spl72_279)), inference(forward_demodulation, [], [f5949, f1022])).
fof(f5949, plain, (~ (op2(e20, e21) = h2(op1(e11, e12))) | (~ spl72_215 | spl72_279)), inference(forward_demodulation, [], [f2591, f2203])).
fof(f2591, plain, (~ (h2(op1(e11, e12)) = op2(h2(e11), e21)) | spl72_279), inference(avatar_component_clause, [], [f2589])).
fof(f2589, plain, (spl72_279 <=> (h2(op1(e11, e12)) = op2(h2(e11), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_279])])).
fof(f5948, plain, (~ spl72_63 | ~ spl72_86 | ~ spl72_208 | spl72_265), inference(avatar_contradiction_clause, [], [f5947])).
fof(f5947, plain, ($false | (~ spl72_63 | ~ spl72_86 | ~ spl72_208 | spl72_265)), inference(subsumption_resolution, [], [f5946, f1254])).
fof(f1254, plain, ((e21 = op2(e22, e22)) | ~ spl72_86), inference(avatar_component_clause, [], [f1252])).
fof(f1252, plain, (spl72_86 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_86])])).
fof(f5946, plain, (~ (e21 = op2(e22, e22)) | (~ spl72_63 | ~ spl72_208 | spl72_265)), inference(forward_demodulation, [], [f5945, f750])).
fof(f5945, plain, (~ (op2(e22, e22) = h2(e12)) | (~ spl72_63 | ~ spl72_208 | spl72_265)), inference(forward_demodulation, [], [f5944, f1124])).
fof(f5944, plain, (~ (op2(e22, e22) = h2(op1(e10, e10))) | (~ spl72_208 | spl72_265)), inference(forward_demodulation, [], [f2535, f2167])).
fof(f2535, plain, (~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10))) | spl72_265), inference(avatar_component_clause, [], [f2533])).
fof(f2533, plain, (spl72_265 <=> (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_265])])).
fof(f5939, plain, (~ spl72_41 | spl72_280), inference(avatar_contradiction_clause, [], [f5938])).
fof(f5938, plain, ($false | (~ spl72_41 | spl72_280)), inference(trivial_inequality_removal, [], [f5937])).
fof(f5937, plain, (~ (h2(e10) = h2(e10)) | (~ spl72_41 | spl72_280)), inference(forward_demodulation, [], [f2595, f1031])).
fof(f1031, plain, ((e10 = op1(e11, e11)) | ~ spl72_41), inference(avatar_component_clause, [], [f1029])).
fof(f1029, plain, (spl72_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_41])])).
fof(f2595, plain, (~ (h2(e10) = h2(op1(e11, e11))) | spl72_280), inference(avatar_component_clause, [], [f2593])).
fof(f2593, plain, (spl72_280 <=> (h2(e10) = h2(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_280])])).
fof(f5918, plain, (~ spl72_56 | ~ spl72_92 | ~ spl72_208 | ~ spl72_273 | spl72_281), inference(avatar_contradiction_clause, [], [f5917])).
fof(f5917, plain, ($false | (~ spl72_56 | ~ spl72_92 | ~ spl72_208 | ~ spl72_273 | spl72_281)), inference(subsumption_resolution, [], [f5916, f1279])).
fof(f1279, plain, ((e23 = op2(e22, e21)) | ~ spl72_92), inference(avatar_component_clause, [], [f1277])).
fof(f1277, plain, (spl72_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_92])])).
fof(f5916, plain, (~ (e23 = op2(e22, e21)) | (~ spl72_56 | ~ spl72_208 | ~ spl72_273 | spl72_281)), inference(forward_demodulation, [], [f5915, f2566])).
fof(f5915, plain, (~ (op2(e22, e21) = h2(e13)) | (~ spl72_56 | ~ spl72_208 | spl72_281)), inference(forward_demodulation, [], [f5914, f1094])).
fof(f5914, plain, (~ (op2(e22, e21) = h2(op1(e10, e12))) | (~ spl72_208 | spl72_281)), inference(forward_demodulation, [], [f2599, f2167])).
fof(f2599, plain, (~ (h2(op1(e10, e12)) = op2(h2(e10), e21)) | spl72_281), inference(avatar_component_clause, [], [f2597])).
fof(f2597, plain, (spl72_281 <=> (h2(op1(e10, e12)) = op2(h2(e10), e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_281])])).
fof(f5905, plain, (~ spl72_1 | ~ spl72_5), inference(avatar_split_clause, [], [f5898, f876, f859])).
fof(f859, plain, (spl72_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_1])])).
fof(f5898, plain, (~ (e10 = op1(e13, e13)) | ~ spl72_5), inference(backward_demodulation, [], [f311, f878])).
fof(f311, plain, ~ (op1(e13, e12) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax5)).
fof(f5891, plain, (~ spl72_8 | ~ spl72_16), inference(avatar_split_clause, [], [f5889, f922, f888])).
fof(f888, plain, (spl72_8 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_8])])).
fof(f5889, plain, (~ (e13 = op1(e13, e12)) | ~ spl72_16), inference(backward_demodulation, [], [f307, f924])).
fof(f307, plain, ~ (op1(e13, e10) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f5879, plain, (~ spl72_34 | ~ spl72_46), inference(avatar_split_clause, [], [f5873, f1050, f999])).
fof(f999, plain, (spl72_34 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_34])])).
fof(f5873, plain, (~ (e11 = op1(e11, e13)) | ~ spl72_46), inference(backward_demodulation, [], [f296, f1052])).
fof(f296, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f5867, plain, (~ spl72_49 | ~ spl72_83 | ~ spl72_208 | spl72_267 | ~ spl72_273), inference(avatar_contradiction_clause, [], [f5866])).
fof(f5866, plain, ($false | (~ spl72_49 | ~ spl72_83 | ~ spl72_208 | spl72_267 | ~ spl72_273)), inference(subsumption_resolution, [], [f5861, f2167])).
fof(f5861, plain, (~ (e22 = h2(e10)) | (~ spl72_49 | ~ spl72_83 | ~ spl72_208 | spl72_267 | ~ spl72_273)), inference(backward_demodulation, [], [f5568, f1065])).
fof(f5568, plain, (~ (e22 = h2(op1(e10, e13))) | (~ spl72_83 | ~ spl72_208 | spl72_267 | ~ spl72_273)), inference(forward_demodulation, [], [f5567, f1241])).
fof(f1241, plain, ((e22 = op2(e22, e23)) | ~ spl72_83), inference(avatar_component_clause, [], [f1239])).
fof(f5567, plain, (~ (op2(e22, e23) = h2(op1(e10, e13))) | (~ spl72_208 | spl72_267 | ~ spl72_273)), inference(forward_demodulation, [], [f5229, f2167])).
fof(f5229, plain, (~ (h2(op1(e10, e13)) = op2(h2(e10), e23)) | (spl72_267 | ~ spl72_273)), inference(forward_demodulation, [], [f2543, f2566])).
fof(f2543, plain, (~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | spl72_267), inference(avatar_component_clause, [], [f2541])).
fof(f2541, plain, (spl72_267 <=> (h2(op1(e10, e13)) = op2(h2(e10), h2(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_267])])).
fof(f5846, plain, (~ spl72_8 | ~ spl72_56), inference(avatar_split_clause, [], [f5841, f1092, f888])).
fof(f5841, plain, (~ (e13 = op1(e13, e12)) | ~ spl72_56), inference(backward_demodulation, [], [f278, f1094])).
fof(f278, plain, ~ (op1(e10, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f5840, plain, (~ spl72_10 | ~ spl72_58), inference(avatar_split_clause, [], [f5832, f1101, f897])).
fof(f897, plain, (spl72_10 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_10])])).
fof(f5832, plain, (~ (e11 = op1(e13, e11)) | ~ spl72_58), inference(backward_demodulation, [], [f272, f1103])).
fof(f272, plain, ~ (op1(e10, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f5831, plain, (~ spl72_59 | ~ spl72_63), inference(avatar_split_clause, [], [f5823, f1122, f1105])).
fof(f1105, plain, (spl72_59 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_59])])).
fof(f5823, plain, (~ (e12 = op1(e10, e11)) | ~ spl72_63), inference(backward_demodulation, [], [f288, f1124])).
fof(f288, plain, ~ (op1(e10, e10) = op1(e10, e11)), inference(cnf_transformation, [], [f5])).
fof(f5830, plain, (~ spl72_15 | ~ spl72_63), inference(avatar_split_clause, [], [f5822, f1122, f918])).
fof(f918, plain, (spl72_15 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_15])])).
fof(f5822, plain, (~ (e12 = op1(e13, e10)) | ~ spl72_63), inference(backward_demodulation, [], [f266, f1124])).
fof(f266, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f5759, plain, (~ spl72_22 | ~ spl72_23), inference(avatar_contradiction_clause, [], [f5758])).
fof(f5758, plain, ($false | (~ spl72_22 | ~ spl72_23)), inference(subsumption_resolution, [], [f5757, f363])).
fof(f363, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax7)).
fof(f5757, plain, ((e11 = e12) | (~ spl72_22 | ~ spl72_23)), inference(forward_demodulation, [], [f954, f950])).
fof(f954, plain, ((e12 = op1(e12, e12)) | ~ spl72_23), inference(avatar_component_clause, [], [f952])).
fof(f952, plain, (spl72_23 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_23])])).
fof(f5741, plain, (~ spl72_62 | ~ spl72_105 | ~ spl72_224 | spl72_227 | ~ spl72_284), inference(avatar_contradiction_clause, [], [f5740])).
fof(f5740, plain, ($false | (~ spl72_62 | ~ spl72_105 | ~ spl72_224 | spl72_227 | ~ spl72_284)), inference(subsumption_resolution, [], [f5736, f2265])).
fof(f2265, plain, (~ (e20 = h1(e11)) | spl72_227), inference(avatar_component_clause, [], [f2263])).
fof(f2263, plain, (spl72_227 <=> (e20 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_227])])).
fof(f5736, plain, ((e20 = h1(e11)) | (~ spl72_62 | ~ spl72_105 | ~ spl72_224 | ~ spl72_284)), inference(backward_demodulation, [], [f5569, f1120])).
fof(f1120, plain, ((op1(e10, e10) = e11) | ~ spl72_62), inference(avatar_component_clause, [], [f1118])).
fof(f1118, plain, (spl72_62 <=> (op1(e10, e10) = e11)), introduced(avatar_definition, [new_symbols(naming, [spl72_62])])).
fof(f5569, plain, ((e20 = h1(op1(e10, e10))) | (~ spl72_105 | ~ spl72_224 | ~ spl72_284)), inference(forward_demodulation, [], [f5444, f1335])).
fof(f1335, plain, ((e20 = op2(e21, e21)) | ~ spl72_105), inference(avatar_component_clause, [], [f1333])).
fof(f1333, plain, (spl72_105 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_105])])).
fof(f5444, plain, ((op2(e21, e21) = h1(op1(e10, e10))) | (~ spl72_224 | ~ spl72_284)), inference(forward_demodulation, [], [f2644, f2248])).
fof(f2644, plain, ((h1(op1(e10, e10)) = op2(h1(e10), h1(e10))) | ~ spl72_284), inference(avatar_component_clause, [], [f2643])).
fof(f2643, plain, (spl72_284 <=> (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_284])])).
fof(f5739, plain, (~ spl72_58 | ~ spl72_62), inference(avatar_split_clause, [], [f5733, f1118, f1101])).
fof(f5733, plain, (~ (e11 = op1(e10, e11)) | ~ spl72_62), inference(backward_demodulation, [], [f288, f1120])).
fof(f5676, plain, (~ spl72_1 | ~ spl72_199 | spl72_248), inference(avatar_contradiction_clause, [], [f5675])).
fof(f5675, plain, ($false | (~ spl72_1 | ~ spl72_199 | spl72_248)), inference(subsumption_resolution, [], [f5669, f2054])).
fof(f5669, plain, (~ (e20 = h3(e10)) | (~ spl72_1 | ~ spl72_199 | spl72_248)), inference(backward_demodulation, [], [f5545, f861])).
fof(f861, plain, ((e10 = op1(e13, e13)) | ~ spl72_1), inference(avatar_component_clause, [], [f859])).
fof(f5545, plain, (~ (e20 = h3(op1(e13, e13))) | (~ spl72_199 | spl72_248)), inference(forward_demodulation, [], [f2424, f2115])).
fof(f2115, plain, ((e20 = h4(e11)) | ~ spl72_199), inference(avatar_component_clause, [], [f2114])).
fof(f2114, plain, (spl72_199 <=> (e20 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_199])])).
fof(f2424, plain, (~ (h4(e11) = h3(op1(e13, e13))) | spl72_248), inference(avatar_component_clause, [], [f2422])).
fof(f2422, plain, (spl72_248 <=> (h4(e11) = h3(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_248])])).
fof(f5645, plain, (~ spl72_29 | ~ spl72_103 | ~ spl72_208 | spl72_278), inference(avatar_contradiction_clause, [], [f5644])).
fof(f5644, plain, ($false | (~ spl72_29 | ~ spl72_103 | ~ spl72_208 | spl72_278)), inference(subsumption_resolution, [], [f5636, f2167])).
fof(f5636, plain, (~ (e22 = h2(e10)) | (~ spl72_29 | ~ spl72_103 | ~ spl72_208 | spl72_278)), inference(backward_demodulation, [], [f5555, f980])).
fof(f5555, plain, (~ (e22 = h2(op1(e12, e10))) | (~ spl72_103 | ~ spl72_208 | spl72_278)), inference(forward_demodulation, [], [f5554, f1326])).
fof(f5554, plain, (~ (op2(e21, e22) = h2(op1(e12, e10))) | (~ spl72_208 | spl72_278)), inference(forward_demodulation, [], [f2587, f2167])).
fof(f2587, plain, (~ (h2(op1(e12, e10)) = op2(e21, h2(e10))) | spl72_278), inference(avatar_component_clause, [], [f2585])).
fof(f2585, plain, (spl72_278 <=> (h2(op1(e12, e10)) = op2(e21, h2(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_278])])).
fof(f5640, plain, (~ spl72_17 | ~ spl72_29), inference(avatar_split_clause, [], [f5631, f978, f927])).
fof(f927, plain, (spl72_17 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_17])])).
fof(f5631, plain, (~ (e10 = op1(e12, e13)) | ~ spl72_29), inference(backward_demodulation, [], [f302, f980])).
fof(f302, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f5639, plain, (~ spl72_13 | ~ spl72_29), inference(avatar_split_clause, [], [f5630, f978, f910])).
fof(f910, plain, (spl72_13 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_13])])).
fof(f5630, plain, (~ (e10 = op1(e13, e10)) | ~ spl72_29), inference(backward_demodulation, [], [f269, f980])).
fof(f269, plain, ~ (op1(e12, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f5629, plain, (~ spl72_34 | ~ spl72_98 | ~ spl72_202 | spl72_256), inference(avatar_contradiction_clause, [], [f5628])).
fof(f5628, plain, ($false | (~ spl72_34 | ~ spl72_98 | ~ spl72_202 | spl72_256)), inference(subsumption_resolution, [], [f5624, f2133])).
fof(f5624, plain, (~ (e21 = h3(e11)) | (~ spl72_34 | ~ spl72_98 | ~ spl72_202 | spl72_256)), inference(backward_demodulation, [], [f5558, f1001])).
fof(f1001, plain, ((e11 = op1(e11, e13)) | ~ spl72_34), inference(avatar_component_clause, [], [f999])).
fof(f5558, plain, (~ (e21 = h3(op1(e11, e13))) | (~ spl72_98 | ~ spl72_202 | spl72_256)), inference(forward_demodulation, [], [f5276, f1305])).
fof(f5276, plain, (~ (op2(e21, e23) = h3(op1(e11, e13))) | (~ spl72_202 | spl72_256)), inference(forward_demodulation, [], [f2456, f2133])).
fof(f2456, plain, (~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | spl72_256), inference(avatar_component_clause, [], [f2454])).
fof(f2454, plain, (spl72_256 <=> (h3(op1(e11, e13)) = op2(h3(e11), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_256])])).
fof(f5625, plain, (~ spl72_2 | ~ spl72_34), inference(avatar_split_clause, [], [f5620, f999, f863])).
fof(f5620, plain, (~ (e11 = op1(e13, e13)) | ~ spl72_34), inference(backward_demodulation, [], [f286, f1001])).
fof(f286, plain, ~ (op1(e11, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f5617, plain, (~ spl72_48 | ~ spl72_112 | ~ spl72_202 | spl72_259), inference(avatar_contradiction_clause, [], [f5616])).
fof(f5616, plain, ($false | (~ spl72_48 | ~ spl72_112 | ~ spl72_202 | spl72_259)), inference(subsumption_resolution, [], [f5612, f2044])).
fof(f2044, plain, (e23 = h3(e13)), inference(forward_demodulation, [], [f757, f745])).
fof(f745, plain, (e23 = op2(e22, op2(e22, e22))), inference(cnf_transformation, [], [f13])).
fof(f757, plain, (op2(e22, op2(e22, e22)) = h3(e13)), inference(cnf_transformation, [], [f16])).
fof(f5612, plain, (~ (e23 = h3(e13)) | (~ spl72_48 | ~ spl72_112 | ~ spl72_202 | spl72_259)), inference(backward_demodulation, [], [f5562, f1060])).
fof(f1060, plain, ((e13 = op1(e11, e10)) | ~ spl72_48), inference(avatar_component_clause, [], [f1058])).
fof(f1058, plain, (spl72_48 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_48])])).
fof(f5562, plain, (~ (e23 = h3(op1(e11, e10))) | (~ spl72_112 | ~ spl72_202 | spl72_259)), inference(forward_demodulation, [], [f5273, f1364])).
fof(f1364, plain, ((e23 = op2(e21, e20)) | ~ spl72_112), inference(avatar_component_clause, [], [f1362])).
fof(f5273, plain, (~ (op2(e21, e20) = h3(op1(e11, e10))) | (~ spl72_202 | spl72_259)), inference(forward_demodulation, [], [f2468, f2133])).
fof(f2468, plain, (~ (h3(op1(e11, e10)) = op2(h3(e11), e20)) | spl72_259), inference(avatar_component_clause, [], [f2466])).
fof(f2466, plain, (spl72_259 <=> (h3(op1(e11, e10)) = op2(h3(e11), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_259])])).
fof(f5614, plain, (~ spl72_36 | ~ spl72_48), inference(avatar_split_clause, [], [f5610, f1058, f1007])).
fof(f5610, plain, (~ (e13 = op1(e11, e13)) | ~ spl72_48), inference(backward_demodulation, [], [f296, f1060])).
fof(f5603, plain, (~ spl72_36 | ~ spl72_52), inference(avatar_split_clause, [], [f5596, f1075, f1007])).
fof(f1075, plain, (spl72_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_52])])).
fof(f5596, plain, (~ (e13 = op1(e11, e13)) | ~ spl72_52), inference(backward_demodulation, [], [f282, f1077])).
fof(f1077, plain, ((e13 = op1(e10, e13)) | ~ spl72_52), inference(avatar_component_clause, [], [f1075])).
fof(f282, plain, ~ (op1(e10, e13) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f5592, plain, (~ spl72_51 | ~ spl72_63), inference(avatar_split_clause, [], [f5586, f1122, f1071])).
fof(f1071, plain, (spl72_51 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_51])])).
fof(f5586, plain, (~ (e12 = op1(e10, e13)) | ~ spl72_63), inference(backward_demodulation, [], [f290, f1124])).
fof(f290, plain, ~ (op1(e10, e10) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f5591, plain, (~ spl72_31 | ~ spl72_63), inference(avatar_split_clause, [], [f5584, f1122, f986])).
fof(f986, plain, (spl72_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_31])])).
fof(f5584, plain, (~ (e12 = op1(e12, e10)) | ~ spl72_63), inference(backward_demodulation, [], [f265, f1124])).
fof(f265, plain, ~ (op1(e10, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f5547, plain, (~ spl72_4 | ~ spl72_8), inference(avatar_split_clause, [], [f5546, f888, f871])).
fof(f871, plain, (spl72_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_4])])).
fof(f5546, plain, (~ (e13 = op1(e13, e13)) | ~ spl72_8), inference(forward_demodulation, [], [f311, f890])).
fof(f890, plain, ((e13 = op1(e13, e12)) | ~ spl72_8), inference(avatar_component_clause, [], [f888])).
fof(f5537, plain, (~ spl72_36 | ~ spl72_122 | ~ spl72_194 | ~ spl72_199 | spl72_233), inference(avatar_contradiction_clause, [], [f5536])).
fof(f5536, plain, ($false | (~ spl72_36 | ~ spl72_122 | ~ spl72_194 | ~ spl72_199 | spl72_233)), inference(subsumption_resolution, [], [f5535, f1407])).
fof(f5535, plain, (~ (e21 = op2(e20, e21)) | (~ spl72_36 | ~ spl72_194 | ~ spl72_199 | spl72_233)), inference(forward_demodulation, [], [f5534, f2089])).
fof(f2089, plain, ((e21 = h4(e13)) | ~ spl72_194), inference(avatar_component_clause, [], [f2088])).
fof(f2088, plain, (spl72_194 <=> (e21 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_194])])).
fof(f5534, plain, (~ (op2(e20, e21) = h4(e13)) | (~ spl72_36 | ~ spl72_194 | ~ spl72_199 | spl72_233)), inference(forward_demodulation, [], [f5533, f1009])).
fof(f5533, plain, (~ (op2(e20, e21) = h4(op1(e11, e13))) | (~ spl72_194 | ~ spl72_199 | spl72_233)), inference(forward_demodulation, [], [f5532, f2115])).
fof(f5532, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), e21)) | (~ spl72_194 | spl72_233)), inference(forward_demodulation, [], [f2301, f2089])).
fof(f2301, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | spl72_233), inference(avatar_component_clause, [], [f2299])).
fof(f2299, plain, (spl72_233 <=> (h4(op1(e11, e13)) = op2(h4(e11), h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_233])])).
fof(f5531, plain, (~ spl72_51 | ~ spl72_92 | ~ spl72_192 | ~ spl72_194 | spl72_231), inference(avatar_contradiction_clause, [], [f5530])).
fof(f5530, plain, ($false | (~ spl72_51 | ~ spl72_92 | ~ spl72_192 | ~ spl72_194 | spl72_231)), inference(subsumption_resolution, [], [f5529, f1279])).
fof(f5529, plain, (~ (e23 = op2(e22, e21)) | (~ spl72_51 | ~ spl72_192 | ~ spl72_194 | spl72_231)), inference(forward_demodulation, [], [f5528, f758])).
fof(f5528, plain, (~ (op2(e22, e21) = h4(e12)) | (~ spl72_51 | ~ spl72_192 | ~ spl72_194 | spl72_231)), inference(forward_demodulation, [], [f5527, f1073])).
fof(f1073, plain, ((e12 = op1(e10, e13)) | ~ spl72_51), inference(avatar_component_clause, [], [f1071])).
fof(f5527, plain, (~ (op2(e22, e21) = h4(op1(e10, e13))) | (~ spl72_192 | ~ spl72_194 | spl72_231)), inference(forward_demodulation, [], [f5526, f2080])).
fof(f2080, plain, ((e22 = h4(e10)) | ~ spl72_192), inference(avatar_component_clause, [], [f2079])).
fof(f2079, plain, (spl72_192 <=> (e22 = h4(e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_192])])).
fof(f5526, plain, (~ (h4(op1(e10, e13)) = op2(h4(e10), e21)) | (~ spl72_194 | spl72_231)), inference(forward_demodulation, [], [f2293, f2089])).
fof(f5519, plain, (~ spl72_8 | ~ spl72_98 | ~ spl72_194 | spl72_238), inference(avatar_contradiction_clause, [], [f5518])).
fof(f5518, plain, ($false | (~ spl72_8 | ~ spl72_98 | ~ spl72_194 | spl72_238)), inference(subsumption_resolution, [], [f5517, f1305])).
fof(f5517, plain, (~ (e21 = op2(e21, e23)) | (~ spl72_8 | ~ spl72_194 | spl72_238)), inference(forward_demodulation, [], [f5516, f2089])).
fof(f5516, plain, (~ (op2(e21, e23) = h4(e13)) | (~ spl72_8 | ~ spl72_194 | spl72_238)), inference(forward_demodulation, [], [f5515, f890])).
fof(f5515, plain, (~ (op2(e21, e23) = h4(op1(e13, e12))) | (~ spl72_194 | spl72_238)), inference(forward_demodulation, [], [f2321, f2089])).
fof(f5514, plain, (~ spl72_2 | ~ spl72_105 | ~ spl72_194 | ~ spl72_199 | spl72_236), inference(avatar_contradiction_clause, [], [f5513])).
fof(f5513, plain, ($false | (~ spl72_2 | ~ spl72_105 | ~ spl72_194 | ~ spl72_199 | spl72_236)), inference(subsumption_resolution, [], [f5512, f1335])).
fof(f5512, plain, (~ (e20 = op2(e21, e21)) | (~ spl72_2 | ~ spl72_194 | ~ spl72_199 | spl72_236)), inference(forward_demodulation, [], [f5511, f2115])).
fof(f5511, plain, (~ (op2(e21, e21) = h4(e11)) | (~ spl72_2 | ~ spl72_194 | spl72_236)), inference(forward_demodulation, [], [f5510, f865])).
fof(f5510, plain, (~ (op2(e21, e21) = h4(op1(e13, e13))) | (~ spl72_194 | spl72_236)), inference(forward_demodulation, [], [f2313, f2089])).
fof(f2313, plain, (~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | spl72_236), inference(avatar_component_clause, [], [f2311])).
fof(f2311, plain, (spl72_236 <=> (h4(op1(e13, e13)) = op2(h4(e13), h4(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_236])])).
fof(f5502, plain, (~ spl72_31 | ~ spl72_72 | ~ spl72_192 | spl72_242), inference(avatar_contradiction_clause, [], [f5501])).
fof(f5501, plain, ($false | (~ spl72_31 | ~ spl72_72 | ~ spl72_192 | spl72_242)), inference(subsumption_resolution, [], [f5500, f1194])).
fof(f5500, plain, (~ (e23 = op2(e23, e22)) | (~ spl72_31 | ~ spl72_192 | spl72_242)), inference(forward_demodulation, [], [f5499, f758])).
fof(f5499, plain, (~ (op2(e23, e22) = h4(e12)) | (~ spl72_31 | ~ spl72_192 | spl72_242)), inference(forward_demodulation, [], [f5498, f988])).
fof(f988, plain, ((e12 = op1(e12, e10)) | ~ spl72_31), inference(avatar_component_clause, [], [f986])).
fof(f5498, plain, (~ (op2(e23, e22) = h4(op1(e12, e10))) | (~ spl72_192 | spl72_242)), inference(forward_demodulation, [], [f2337, f2080])).
fof(f5493, plain, (~ spl72_11 | ~ spl72_112 | ~ spl72_194 | ~ spl72_199 | spl72_235), inference(avatar_contradiction_clause, [], [f5492])).
fof(f5492, plain, ($false | (~ spl72_11 | ~ spl72_112 | ~ spl72_194 | ~ spl72_199 | spl72_235)), inference(subsumption_resolution, [], [f5491, f1364])).
fof(f5491, plain, (~ (e23 = op2(e21, e20)) | (~ spl72_11 | ~ spl72_194 | ~ spl72_199 | spl72_235)), inference(forward_demodulation, [], [f5490, f758])).
fof(f5490, plain, (~ (op2(e21, e20) = h4(e12)) | (~ spl72_11 | ~ spl72_194 | ~ spl72_199 | spl72_235)), inference(forward_demodulation, [], [f5489, f903])).
fof(f903, plain, ((e12 = op1(e13, e11)) | ~ spl72_11), inference(avatar_component_clause, [], [f901])).
fof(f901, plain, (spl72_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_11])])).
fof(f5489, plain, (~ (op2(e21, e20) = h4(op1(e13, e11))) | (~ spl72_194 | ~ spl72_199 | spl72_235)), inference(forward_demodulation, [], [f5488, f2089])).
fof(f5488, plain, (~ (h4(op1(e13, e11)) = op2(h4(e13), e20)) | (~ spl72_199 | spl72_235)), inference(forward_demodulation, [], [f2309, f2115])).
fof(f2309, plain, (~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | spl72_235), inference(avatar_component_clause, [], [f2307])).
fof(f2307, plain, (spl72_235 <=> (h4(op1(e13, e11)) = op2(h4(e13), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_235])])).
fof(f5483, plain, (~ spl72_13 | ~ spl72_103 | ~ spl72_192 | ~ spl72_194 | spl72_234), inference(avatar_contradiction_clause, [], [f5482])).
fof(f5482, plain, ($false | (~ spl72_13 | ~ spl72_103 | ~ spl72_192 | ~ spl72_194 | spl72_234)), inference(subsumption_resolution, [], [f5481, f1326])).
fof(f5481, plain, (~ (e22 = op2(e21, e22)) | (~ spl72_13 | ~ spl72_192 | ~ spl72_194 | spl72_234)), inference(forward_demodulation, [], [f5480, f2080])).
fof(f5480, plain, (~ (op2(e21, e22) = h4(e10)) | (~ spl72_13 | ~ spl72_192 | ~ spl72_194 | spl72_234)), inference(forward_demodulation, [], [f5479, f912])).
fof(f912, plain, ((e10 = op1(e13, e10)) | ~ spl72_13), inference(avatar_component_clause, [], [f910])).
fof(f5479, plain, (~ (op2(e21, e22) = h4(op1(e13, e10))) | (~ spl72_192 | ~ spl72_194 | spl72_234)), inference(forward_demodulation, [], [f5478, f2089])).
fof(f5478, plain, (~ (h4(op1(e13, e10)) = op2(h4(e13), e22)) | (~ spl72_192 | spl72_234)), inference(forward_demodulation, [], [f2305, f2080])).
fof(f5477, plain, (~ spl72_58 | ~ spl72_93 | ~ spl72_192 | ~ spl72_199 | spl72_230), inference(avatar_contradiction_clause, [], [f5476])).
fof(f5476, plain, ($false | (~ spl72_58 | ~ spl72_93 | ~ spl72_192 | ~ spl72_199 | spl72_230)), inference(subsumption_resolution, [], [f5475, f1284])).
fof(f5475, plain, (~ (e20 = op2(e22, e20)) | (~ spl72_58 | ~ spl72_192 | ~ spl72_199 | spl72_230)), inference(forward_demodulation, [], [f5474, f2115])).
fof(f5474, plain, (~ (op2(e22, e20) = h4(e11)) | (~ spl72_58 | ~ spl72_192 | ~ spl72_199 | spl72_230)), inference(forward_demodulation, [], [f5473, f1103])).
fof(f5473, plain, (~ (op2(e22, e20) = h4(op1(e10, e11))) | (~ spl72_192 | ~ spl72_199 | spl72_230)), inference(forward_demodulation, [], [f5472, f2080])).
fof(f5472, plain, (~ (h4(op1(e10, e11)) = op2(h4(e10), e20)) | (~ spl72_199 | spl72_230)), inference(forward_demodulation, [], [f2289, f2115])).
fof(f2289, plain, (~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | spl72_230), inference(avatar_component_clause, [], [f2287])).
fof(f2287, plain, (spl72_230 <=> (h4(op1(e10, e11)) = op2(h4(e10), h4(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_230])])).
fof(f5452, plain, (~ spl72_53 | ~ spl72_83 | ~ spl72_192 | spl72_245), inference(avatar_contradiction_clause, [], [f5451])).
fof(f5451, plain, ($false | (~ spl72_53 | ~ spl72_83 | ~ spl72_192 | spl72_245)), inference(subsumption_resolution, [], [f5450, f1241])).
fof(f5450, plain, (~ (e22 = op2(e22, e23)) | (~ spl72_53 | ~ spl72_192 | spl72_245)), inference(forward_demodulation, [], [f5449, f2080])).
fof(f5449, plain, (~ (op2(e22, e23) = h4(e10)) | (~ spl72_53 | ~ spl72_192 | spl72_245)), inference(forward_demodulation, [], [f5448, f1082])).
fof(f1082, plain, ((e10 = op1(e10, e12)) | ~ spl72_53), inference(avatar_component_clause, [], [f1080])).
fof(f1080, plain, (spl72_53 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_53])])).
fof(f5448, plain, (~ (op2(e22, e23) = h4(op1(e10, e12))) | (~ spl72_192 | spl72_245)), inference(forward_demodulation, [], [f2349, f2080])).
fof(f5438, plain, (~ spl72_46 | ~ spl72_117 | ~ spl72_192 | ~ spl72_199 | spl72_232), inference(avatar_contradiction_clause, [], [f5437])).
fof(f5437, plain, ($false | (~ spl72_46 | ~ spl72_117 | ~ spl72_192 | ~ spl72_199 | spl72_232)), inference(subsumption_resolution, [], [f5436, f1386])).
fof(f5436, plain, (~ (e20 = op2(e20, e22)) | (~ spl72_46 | ~ spl72_192 | ~ spl72_199 | spl72_232)), inference(forward_demodulation, [], [f5435, f2115])).
fof(f5435, plain, (~ (op2(e20, e22) = h4(e11)) | (~ spl72_46 | ~ spl72_192 | ~ spl72_199 | spl72_232)), inference(forward_demodulation, [], [f5434, f1052])).
fof(f5434, plain, (~ (op2(e20, e22) = h4(op1(e11, e10))) | (~ spl72_192 | ~ spl72_199 | spl72_232)), inference(forward_demodulation, [], [f5433, f2115])).
fof(f5433, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), e22)) | (~ spl72_192 | spl72_232)), inference(forward_demodulation, [], [f2297, f2080])).
fof(f2297, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | spl72_232), inference(avatar_component_clause, [], [f2295])).
fof(f2295, plain, (spl72_232 <=> (h4(op1(e11, e10)) = op2(h4(e11), h4(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_232])])).
fof(f5429, plain, (spl72_226 | ~ spl72_117 | ~ spl72_219), inference(avatar_split_clause, [], [f5428, f2222, f1384, f2256])).
fof(f2256, plain, (spl72_226 <=> (e20 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_226])])).
fof(f5428, plain, ((e20 = h1(e13)) | (~ spl72_117 | ~ spl72_219)), inference(forward_demodulation, [], [f5427, f1386])).
fof(f5427, plain, ((op2(e20, e22) = h1(e13)) | ~ spl72_219), inference(forward_demodulation, [], [f2034, f2223])).
fof(f2223, plain, ((e22 = h1(e11)) | ~ spl72_219), inference(avatar_component_clause, [], [f2222])).
fof(f2034, plain, (h1(e13) = op2(e20, h1(e11))), inference(backward_demodulation, [], [f749, f748])).
fof(f749, plain, (h1(e13) = op2(e20, op2(e20, e20))), inference(cnf_transformation, [], [f14])).
fof(f5425, plain, (~ spl72_192 | ~ spl72_17 | ~ spl72_75 | ~ spl72_194 | spl72_239), inference(avatar_split_clause, [], [f5424, f2323, f2088, f1205, f927, f2079])).
fof(f5424, plain, (~ (e22 = h4(e10)) | (~ spl72_17 | ~ spl72_75 | ~ spl72_194 | spl72_239)), inference(forward_demodulation, [], [f5422, f1207])).
fof(f5422, plain, (~ (op2(e23, e21) = h4(e10)) | (~ spl72_17 | ~ spl72_194 | spl72_239)), inference(backward_demodulation, [], [f5397, f2089])).
fof(f5397, plain, (~ (h4(e10) = op2(e23, h4(e13))) | (~ spl72_17 | spl72_239)), inference(forward_demodulation, [], [f2325, f929])).
fof(f929, plain, ((e10 = op1(e12, e13)) | ~ spl72_17), inference(avatar_component_clause, [], [f927])).
fof(f5388, plain, (spl72_194 | ~ spl72_78 | ~ spl72_199), inference(avatar_split_clause, [], [f5387, f2114, f1218, f2088])).
fof(f5387, plain, ((e21 = h4(e13)) | (~ spl72_78 | ~ spl72_199)), inference(backward_demodulation, [], [f5201, f1220])).
fof(f1220, plain, ((e21 = op2(e23, e20)) | ~ spl72_78), inference(avatar_component_clause, [], [f1218])).
fof(f5201, plain, ((op2(e23, e20) = h4(e13)) | ~ spl72_199), inference(forward_demodulation, [], [f2061, f2115])).
fof(f2061, plain, (h4(e13) = op2(e23, h4(e11))), inference(backward_demodulation, [], [f761, f760])).
fof(f761, plain, (h4(e13) = op2(e23, op2(e23, e23))), inference(cnf_transformation, [], [f17])).
fof(f5383, plain, (~ spl72_74 | ~ spl72_122), inference(avatar_split_clause, [], [f5382, f1405, f1201])).
fof(f1201, plain, (spl72_74 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_74])])).
fof(f5382, plain, (~ (e21 = op2(e23, e21)) | ~ spl72_122), inference(backward_demodulation, [], [f320, f1407])).
fof(f320, plain, ~ (op2(e20, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f5377, plain, (spl72_192 | ~ spl72_199 | ~ spl72_219), inference(avatar_split_clause, [], [f5363, f2222, f2114, f2079])).
fof(f5363, plain, ((e22 = h4(e10)) | (~ spl72_199 | ~ spl72_219)), inference(backward_demodulation, [], [f5315, f2223])).
fof(f5315, plain, ((h1(e11) = h4(e10)) | ~ spl72_199), inference(backward_demodulation, [], [f5215, f748])).
fof(f5215, plain, ((op2(e20, e20) = h4(e10)) | ~ spl72_199), inference(forward_demodulation, [], [f2062, f2115])).
fof(f2062, plain, (h4(e10) = op2(h4(e11), h4(e11))), inference(forward_demodulation, [], [f759, f760])).
fof(f759, plain, (h4(e10) = op2(op2(e23, e23), op2(e23, e23))), inference(cnf_transformation, [], [f17])).
fof(f5376, plain, (spl72_208 | ~ spl72_215 | ~ spl72_219), inference(avatar_split_clause, [], [f5357, f2222, f2202, f2166])).
fof(f5357, plain, ((e22 = h2(e10)) | (~ spl72_215 | ~ spl72_219)), inference(backward_demodulation, [], [f3351, f2223])).
fof(f3351, plain, ((h1(e11) = h2(e10)) | ~ spl72_215), inference(forward_demodulation, [], [f3342, f748])).
fof(f3342, plain, ((op2(e20, e20) = h2(e10)) | ~ spl72_215), inference(backward_demodulation, [], [f2043, f2203])).
fof(f2043, plain, (h2(e10) = op2(h2(e11), h2(e11))), inference(forward_demodulation, [], [f751, f752])).
fof(f751, plain, (h2(e10) = op2(op2(e21, e21), op2(e21, e21))), inference(cnf_transformation, [], [f15])).
fof(f5375, plain, (~ spl72_123 | ~ spl72_219), inference(avatar_split_clause, [], [f5355, f2222, f1409])).
fof(f1409, plain, (spl72_123 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_123])])).
fof(f5355, plain, (~ (e22 = op2(e20, e21)) | ~ spl72_219), inference(backward_demodulation, [], [f2031, f2223])).
fof(f2031, plain, ~ (op2(e20, e21) = h1(e11)), inference(backward_demodulation, [], [f336, f748])).
fof(f336, plain, ~ (op2(e20, e20) = op2(e20, e21)), inference(cnf_transformation, [], [f6])).
fof(f5374, plain, (~ spl72_79 | ~ spl72_219), inference(avatar_split_clause, [], [f5354, f2222, f1222])).
fof(f1222, plain, (spl72_79 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_79])])).
fof(f5354, plain, (~ (e22 = op2(e23, e20)) | ~ spl72_219), inference(backward_demodulation, [], [f2030, f2223])).
fof(f2030, plain, ~ (op2(e23, e20) = h1(e11)), inference(backward_demodulation, [], [f314, f748])).
fof(f314, plain, ~ (op2(e20, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f5352, plain, (~ spl72_194 | ~ spl72_64 | ~ spl72_199 | ~ spl72_224 | spl72_229), inference(avatar_split_clause, [], [f5349, f2283, f2247, f2114, f1126, f2088])).
fof(f1126, plain, (spl72_64 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl72_64])])).
fof(f5349, plain, (~ (e21 = h4(e13)) | (~ spl72_64 | ~ spl72_199 | ~ spl72_224 | spl72_229)), inference(backward_demodulation, [], [f5326, f2248])).
fof(f5326, plain, (~ (h1(e10) = h4(e13)) | (~ spl72_64 | ~ spl72_199 | spl72_229)), inference(forward_demodulation, [], [f5322, f2035])).
fof(f5322, plain, (~ (h4(e13) = op2(h1(e11), h1(e11))) | (~ spl72_64 | ~ spl72_199 | spl72_229)), inference(backward_demodulation, [], [f5289, f5315])).
fof(f5289, plain, (~ (h4(e13) = op2(h4(e10), h4(e10))) | (~ spl72_64 | spl72_229)), inference(forward_demodulation, [], [f2285, f1128])).
fof(f1128, plain, ((op1(e10, e10) = e13) | ~ spl72_64), inference(avatar_component_clause, [], [f1126])).
fof(f2285, plain, (~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10))) | spl72_229), inference(avatar_component_clause, [], [f2283])).
fof(f5299, plain, (~ spl72_78 | ~ spl72_11 | ~ spl72_215 | spl72_271 | ~ spl72_273), inference(avatar_split_clause, [], [f5269, f2565, f2557, f2202, f901, f1218])).
fof(f2557, plain, (spl72_271 <=> (h2(op1(e13, e11)) = op2(h2(e13), h2(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_271])])).
fof(f5269, plain, (~ (e21 = op2(e23, e20)) | (~ spl72_11 | ~ spl72_215 | spl72_271 | ~ spl72_273)), inference(forward_demodulation, [], [f5268, f750])).
fof(f5268, plain, (~ (op2(e23, e20) = h2(e12)) | (~ spl72_11 | ~ spl72_215 | spl72_271 | ~ spl72_273)), inference(forward_demodulation, [], [f5267, f903])).
fof(f5267, plain, (~ (op2(e23, e20) = h2(op1(e13, e11))) | (~ spl72_215 | spl72_271 | ~ spl72_273)), inference(forward_demodulation, [], [f5266, f2566])).
fof(f5266, plain, (~ (h2(op1(e13, e11)) = op2(h2(e13), e20)) | (~ spl72_215 | spl72_271)), inference(forward_demodulation, [], [f2559, f2203])).
fof(f2559, plain, (~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | spl72_271), inference(avatar_component_clause, [], [f2557])).
fof(f5296, plain, (~ spl72_86 | ~ spl72_87), inference(avatar_contradiction_clause, [], [f5295])).
fof(f5295, plain, ($false | (~ spl72_86 | ~ spl72_87)), inference(subsumption_resolution, [], [f5294, f369])).
fof(f369, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax8)).
fof(f5294, plain, ((e21 = e22) | (~ spl72_86 | ~ spl72_87)), inference(forward_demodulation, [], [f1258, f1254])).
fof(f1258, plain, ((e22 = op2(e22, e22)) | ~ spl72_87), inference(avatar_component_clause, [], [f1256])).
fof(f1256, plain, (spl72_87 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_87])])).
fof(f5194, plain, (~ spl72_75 | ~ spl72_79), inference(avatar_split_clause, [], [f5191, f1222, f1205])).
fof(f5191, plain, (~ (e22 = op2(e23, e21)) | ~ spl72_79), inference(backward_demodulation, [], [f354, f1224])).
fof(f1224, plain, ((e22 = op2(e23, e20)) | ~ spl72_79), inference(avatar_component_clause, [], [f1222])).
fof(f354, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f5187, plain, (~ spl72_105 | ~ spl72_108), inference(avatar_contradiction_clause, [], [f5186])).
fof(f5186, plain, ($false | (~ spl72_105 | ~ spl72_108)), inference(subsumption_resolution, [], [f5185, f368])).
fof(f368, plain, ~ (e20 = e23), inference(cnf_transformation, [], [f8])).
fof(f5185, plain, ((e20 = e23) | (~ spl72_105 | ~ spl72_108)), inference(forward_demodulation, [], [f1347, f1335])).
fof(f1347, plain, ((e23 = op2(e21, e21)) | ~ spl72_108), inference(avatar_component_clause, [], [f1345])).
fof(f1345, plain, (spl72_108 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_108])])).
fof(f5184, plain, (~ spl72_100 | ~ spl72_116), inference(avatar_split_clause, [], [f5183, f1379, f1311])).
fof(f1311, plain, (spl72_100 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_100])])).
fof(f5183, plain, (~ (e23 = op2(e21, e23)) | ~ spl72_116), inference(backward_demodulation, [], [f330, f1381])).
fof(f330, plain, ~ (op2(e20, e23) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f5182, plain, (~ spl72_69 | ~ spl72_117), inference(avatar_split_clause, [], [f5178, f1384, f1180])).
fof(f1180, plain, (spl72_69 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_69])])).
fof(f5178, plain, (~ (e20 = op2(e23, e22)) | ~ spl72_117), inference(backward_demodulation, [], [f326, f1386])).
fof(f326, plain, ~ (op2(e20, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f5181, plain, (~ spl72_113 | ~ spl72_117), inference(avatar_split_clause, [], [f5177, f1384, f1367])).
fof(f1367, plain, (spl72_113 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_113])])).
fof(f5177, plain, (~ (e20 = op2(e20, e23)) | ~ spl72_117), inference(backward_demodulation, [], [f341, f1386])).
fof(f341, plain, ~ (op2(e20, e22) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f5168, plain, (~ spl72_116 | ~ spl72_39 | ~ spl72_199 | spl72_243), inference(avatar_split_clause, [], [f5152, f2339, f2114, f1020, f1379])).
fof(f2339, plain, (spl72_243 <=> (h4(op1(e11, e12)) = op2(h4(e11), e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_243])])).
fof(f5152, plain, (~ (e23 = op2(e20, e23)) | (~ spl72_39 | ~ spl72_199 | spl72_243)), inference(backward_demodulation, [], [f5100, f2115])).
fof(f5100, plain, (~ (e23 = op2(h4(e11), e23)) | (~ spl72_39 | spl72_243)), inference(forward_demodulation, [], [f5099, f758])).
fof(f5099, plain, (~ (h4(e12) = op2(h4(e11), e23)) | (~ spl72_39 | spl72_243)), inference(forward_demodulation, [], [f2341, f1022])).
fof(f2341, plain, (~ (h4(op1(e11, e12)) = op2(h4(e11), e23)) | spl72_243), inference(avatar_component_clause, [], [f2339])).
fof(f5137, plain, (~ spl72_211 | ~ spl72_215), inference(avatar_contradiction_clause, [], [f5136])).
fof(f5136, plain, ($false | (~ spl72_211 | ~ spl72_215)), inference(subsumption_resolution, [], [f5135, f366])).
fof(f366, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f8])).
fof(f5135, plain, ((e20 = e21) | (~ spl72_211 | ~ spl72_215)), inference(forward_demodulation, [], [f2183, f2203])).
fof(f2183, plain, ((e21 = h2(e11)) | ~ spl72_211), inference(avatar_component_clause, [], [f2182])).
fof(f2182, plain, (spl72_211 <=> (e21 = h2(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_211])])).
fof(f5133, plain, (spl72_124 | ~ spl72_223 | ~ spl72_292), inference(avatar_contradiction_clause, [], [f5132])).
fof(f5132, plain, ($false | (spl72_124 | ~ spl72_223 | ~ spl72_292)), inference(subsumption_resolution, [], [f5115, f1414])).
fof(f1414, plain, (~ (e23 = op2(e20, e21)) | spl72_124), inference(avatar_component_clause, [], [f1413])).
fof(f1413, plain, (spl72_124 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_124])])).
fof(f5115, plain, ((e23 = op2(e20, e21)) | (~ spl72_223 | ~ spl72_292)), inference(backward_demodulation, [], [f5081, f2243])).
fof(f2243, plain, ((e21 = h1(e11)) | ~ spl72_223), inference(avatar_component_clause, [], [f2242])).
fof(f2242, plain, (spl72_223 <=> (e21 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_223])])).
fof(f5081, plain, ((e23 = op2(e20, h1(e11))) | ~ spl72_292), inference(forward_demodulation, [], [f2034, f2676])).
fof(f2676, plain, ((e23 = h1(e13)) | ~ spl72_292), inference(avatar_component_clause, [], [f2675])).
fof(f2675, plain, (spl72_292 <=> (e23 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_292])])).
fof(f5118, plain, (~ spl72_122 | ~ spl72_223), inference(avatar_split_clause, [], [f5104, f2242, f1405])).
fof(f5104, plain, (~ (e21 = op2(e20, e21)) | ~ spl72_223), inference(backward_demodulation, [], [f2031, f2243])).
fof(f5117, plain, (~ spl72_110 | ~ spl72_223), inference(avatar_split_clause, [], [f5102, f2242, f1354])).
fof(f1354, plain, (spl72_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_110])])).
fof(f5102, plain, (~ (e21 = op2(e21, e20)) | ~ spl72_223), inference(backward_demodulation, [], [f2028, f2243])).
fof(f5098, plain, (spl72_72 | ~ spl72_8 | ~ spl72_249), inference(avatar_split_clause, [], [f5097, f2426, f888, f1192])).
fof(f2426, plain, (spl72_249 <=> (op2(e23, e22) = h3(op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_249])])).
fof(f5097, plain, ((e23 = op2(e23, e22)) | (~ spl72_8 | ~ spl72_249)), inference(forward_demodulation, [], [f5096, f2044])).
fof(f5096, plain, ((op2(e23, e22) = h3(e13)) | (~ spl72_8 | ~ spl72_249)), inference(forward_demodulation, [], [f2427, f890])).
fof(f2427, plain, ((op2(e23, e22) = h3(op1(e13, e12))) | ~ spl72_249), inference(avatar_component_clause, [], [f2426])).
fof(f5089, plain, (~ spl72_80 | ~ spl72_215 | ~ spl72_273), inference(avatar_split_clause, [], [f5084, f2565, f2202, f1226])).
fof(f1226, plain, (spl72_80 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_80])])).
fof(f5084, plain, (~ (e23 = op2(e23, e20)) | (~ spl72_215 | ~ spl72_273)), inference(backward_demodulation, [], [f3349, f2566])).
fof(f5058, plain, (~ spl72_191 | ~ spl72_83), inference(avatar_split_clause, [], [f5057, f1239, f2074])).
fof(f2074, plain, (spl72_191 <=> (e22 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_191])])).
fof(f5057, plain, (~ (e22 = h4(e11)) | ~ spl72_83), inference(forward_demodulation, [], [f2057, f1241])).
fof(f2057, plain, ~ (op2(e22, e23) = h4(e11)), inference(backward_demodulation, [], [f335, f760])).
fof(f335, plain, ~ (op2(e22, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f5012, plain, (~ spl72_53 | ~ spl72_113 | ~ spl72_200 | spl72_245), inference(avatar_contradiction_clause, [], [f5011])).
fof(f5011, plain, ($false | (~ spl72_53 | ~ spl72_113 | ~ spl72_200 | spl72_245)), inference(subsumption_resolution, [], [f5010, f1369])).
fof(f1369, plain, ((e20 = op2(e20, e23)) | ~ spl72_113), inference(avatar_component_clause, [], [f1367])).
fof(f5010, plain, (~ (e20 = op2(e20, e23)) | (~ spl72_53 | ~ spl72_200 | spl72_245)), inference(forward_demodulation, [], [f5009, f2120])).
fof(f5009, plain, (~ (op2(e20, e23) = h4(e10)) | (~ spl72_53 | ~ spl72_200 | spl72_245)), inference(forward_demodulation, [], [f5008, f1082])).
fof(f5008, plain, (~ (op2(e20, e23) = h4(op1(e10, e12))) | (~ spl72_200 | spl72_245)), inference(forward_demodulation, [], [f2349, f2120])).
fof(f5004, plain, (~ spl72_41 | spl72_244), inference(avatar_contradiction_clause, [], [f5003])).
fof(f5003, plain, ($false | (~ spl72_41 | spl72_244)), inference(trivial_inequality_removal, [], [f5002])).
fof(f5002, plain, (~ (h4(e10) = h4(e10)) | (~ spl72_41 | spl72_244)), inference(forward_demodulation, [], [f2345, f1031])).
fof(f2345, plain, (~ (h4(e10) = h4(op1(e11, e11))) | spl72_244), inference(avatar_component_clause, [], [f2343])).
fof(f2343, plain, (spl72_244 <=> (h4(e10) = h4(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_244])])).
fof(f4998, plain, (~ spl72_39 | ~ spl72_100 | ~ spl72_195 | spl72_243), inference(avatar_contradiction_clause, [], [f4997])).
fof(f4997, plain, ($false | (~ spl72_39 | ~ spl72_100 | ~ spl72_195 | spl72_243)), inference(subsumption_resolution, [], [f4996, f1313])).
fof(f1313, plain, ((e23 = op2(e21, e23)) | ~ spl72_100), inference(avatar_component_clause, [], [f1311])).
fof(f4996, plain, (~ (e23 = op2(e21, e23)) | (~ spl72_39 | ~ spl72_195 | spl72_243)), inference(forward_demodulation, [], [f4995, f758])).
fof(f4995, plain, (~ (op2(e21, e23) = h4(e12)) | (~ spl72_39 | ~ spl72_195 | spl72_243)), inference(forward_demodulation, [], [f4994, f1022])).
fof(f4994, plain, (~ (op2(e21, e23) = h4(op1(e11, e12))) | (~ spl72_195 | spl72_243)), inference(forward_demodulation, [], [f2341, f2095])).
fof(f4993, plain, (~ spl72_31 | ~ spl72_80 | ~ spl72_200 | spl72_242), inference(avatar_contradiction_clause, [], [f4992])).
fof(f4992, plain, ($false | (~ spl72_31 | ~ spl72_80 | ~ spl72_200 | spl72_242)), inference(subsumption_resolution, [], [f4991, f1228])).
fof(f1228, plain, ((e23 = op2(e23, e20)) | ~ spl72_80), inference(avatar_component_clause, [], [f1226])).
fof(f4991, plain, (~ (e23 = op2(e23, e20)) | (~ spl72_31 | ~ spl72_200 | spl72_242)), inference(forward_demodulation, [], [f4990, f758])).
fof(f4990, plain, (~ (op2(e23, e20) = h4(e12)) | (~ spl72_31 | ~ spl72_200 | spl72_242)), inference(forward_demodulation, [], [f4989, f988])).
fof(f4989, plain, (~ (op2(e23, e20) = h4(op1(e12, e10))) | (~ spl72_200 | spl72_242)), inference(forward_demodulation, [], [f2337, f2120])).
fof(f4988, plain, (~ spl72_28 | spl72_241), inference(avatar_contradiction_clause, [], [f4987])).
fof(f4987, plain, ($false | (~ spl72_28 | spl72_241)), inference(trivial_inequality_removal, [], [f4986])).
fof(f4986, plain, (~ (h4(e13) = h4(e13)) | (~ spl72_28 | spl72_241)), inference(forward_demodulation, [], [f2333, f975])).
fof(f2333, plain, (~ (h4(e13) = h4(op1(e12, e11))) | spl72_241), inference(avatar_component_clause, [], [f2331])).
fof(f2331, plain, (spl72_241 <=> (h4(e13) = h4(op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_241])])).
fof(f4982, plain, (~ spl72_22 | spl72_240), inference(avatar_contradiction_clause, [], [f4981])).
fof(f4981, plain, ($false | (~ spl72_22 | spl72_240)), inference(trivial_inequality_removal, [], [f4980])).
fof(f4980, plain, (~ (h4(e11) = h4(e11)) | (~ spl72_22 | spl72_240)), inference(forward_demodulation, [], [f2329, f950])).
fof(f2329, plain, (~ (h4(e11) = h4(op1(e12, e12))) | spl72_240), inference(avatar_component_clause, [], [f2327])).
fof(f2327, plain, (spl72_240 <=> (h4(e11) = h4(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_240])])).
fof(f4976, plain, (~ spl72_17 | ~ spl72_69 | ~ spl72_190 | ~ spl72_200 | spl72_239), inference(avatar_contradiction_clause, [], [f4975])).
fof(f4975, plain, ($false | (~ spl72_17 | ~ spl72_69 | ~ spl72_190 | ~ spl72_200 | spl72_239)), inference(subsumption_resolution, [], [f4974, f1182])).
fof(f1182, plain, ((e20 = op2(e23, e22)) | ~ spl72_69), inference(avatar_component_clause, [], [f1180])).
fof(f4974, plain, (~ (e20 = op2(e23, e22)) | (~ spl72_17 | ~ spl72_190 | ~ spl72_200 | spl72_239)), inference(forward_demodulation, [], [f4973, f2120])).
fof(f4973, plain, (~ (op2(e23, e22) = h4(e10)) | (~ spl72_17 | ~ spl72_190 | spl72_239)), inference(forward_demodulation, [], [f4972, f929])).
fof(f4972, plain, (~ (op2(e23, e22) = h4(op1(e12, e13))) | (~ spl72_190 | spl72_239)), inference(forward_demodulation, [], [f2325, f2069])).
fof(f4968, plain, (~ spl72_8 | ~ spl72_83 | ~ spl72_190 | spl72_238), inference(avatar_contradiction_clause, [], [f4967])).
fof(f4967, plain, ($false | (~ spl72_8 | ~ spl72_83 | ~ spl72_190 | spl72_238)), inference(subsumption_resolution, [], [f4966, f1241])).
fof(f4966, plain, (~ (e22 = op2(e22, e23)) | (~ spl72_8 | ~ spl72_190 | spl72_238)), inference(forward_demodulation, [], [f4965, f2069])).
fof(f4965, plain, (~ (op2(e22, e23) = h4(e13)) | (~ spl72_8 | ~ spl72_190 | spl72_238)), inference(forward_demodulation, [], [f4964, f890])).
fof(f4964, plain, (~ (op2(e22, e23) = h4(op1(e13, e12))) | (~ spl72_190 | spl72_238)), inference(forward_demodulation, [], [f2321, f2069])).
fof(f4959, plain, (~ spl72_2 | ~ spl72_86 | ~ spl72_190 | ~ spl72_195 | spl72_236), inference(avatar_contradiction_clause, [], [f4958])).
fof(f4958, plain, ($false | (~ spl72_2 | ~ spl72_86 | ~ spl72_190 | ~ spl72_195 | spl72_236)), inference(subsumption_resolution, [], [f4957, f1254])).
fof(f4957, plain, (~ (e21 = op2(e22, e22)) | (~ spl72_2 | ~ spl72_190 | ~ spl72_195 | spl72_236)), inference(forward_demodulation, [], [f4956, f2095])).
fof(f4956, plain, (~ (op2(e22, e22) = h4(e11)) | (~ spl72_2 | ~ spl72_190 | spl72_236)), inference(forward_demodulation, [], [f4955, f865])).
fof(f4955, plain, (~ (op2(e22, e22) = h4(op1(e13, e13))) | (~ spl72_190 | spl72_236)), inference(forward_demodulation, [], [f2313, f2069])).
fof(f4950, plain, (~ spl72_11 | ~ spl72_92 | ~ spl72_190 | ~ spl72_195 | spl72_235), inference(avatar_contradiction_clause, [], [f4949])).
fof(f4949, plain, ($false | (~ spl72_11 | ~ spl72_92 | ~ spl72_190 | ~ spl72_195 | spl72_235)), inference(subsumption_resolution, [], [f4948, f1279])).
fof(f4948, plain, (~ (e23 = op2(e22, e21)) | (~ spl72_11 | ~ spl72_190 | ~ spl72_195 | spl72_235)), inference(forward_demodulation, [], [f4947, f758])).
fof(f4947, plain, (~ (op2(e22, e21) = h4(e12)) | (~ spl72_11 | ~ spl72_190 | ~ spl72_195 | spl72_235)), inference(forward_demodulation, [], [f4946, f903])).
fof(f4946, plain, (~ (op2(e22, e21) = h4(op1(e13, e11))) | (~ spl72_190 | ~ spl72_195 | spl72_235)), inference(forward_demodulation, [], [f4945, f2069])).
fof(f4945, plain, (~ (h4(op1(e13, e11)) = op2(h4(e13), e21)) | (~ spl72_195 | spl72_235)), inference(forward_demodulation, [], [f2309, f2095])).
fof(f4940, plain, (~ spl72_13 | ~ spl72_93 | ~ spl72_190 | ~ spl72_200 | spl72_234), inference(avatar_contradiction_clause, [], [f4939])).
fof(f4939, plain, ($false | (~ spl72_13 | ~ spl72_93 | ~ spl72_190 | ~ spl72_200 | spl72_234)), inference(subsumption_resolution, [], [f4938, f1284])).
fof(f4938, plain, (~ (e20 = op2(e22, e20)) | (~ spl72_13 | ~ spl72_190 | ~ spl72_200 | spl72_234)), inference(forward_demodulation, [], [f4937, f2120])).
fof(f4937, plain, (~ (op2(e22, e20) = h4(e10)) | (~ spl72_13 | ~ spl72_190 | ~ spl72_200 | spl72_234)), inference(forward_demodulation, [], [f4936, f912])).
fof(f4936, plain, (~ (op2(e22, e20) = h4(op1(e13, e10))) | (~ spl72_190 | ~ spl72_200 | spl72_234)), inference(forward_demodulation, [], [f4935, f2069])).
fof(f4935, plain, (~ (h4(op1(e13, e10)) = op2(h4(e13), e20)) | (~ spl72_200 | spl72_234)), inference(forward_demodulation, [], [f2305, f2120])).
fof(f4930, plain, (~ spl72_36 | ~ spl72_103 | ~ spl72_190 | ~ spl72_195 | spl72_233), inference(avatar_contradiction_clause, [], [f4929])).
fof(f4929, plain, ($false | (~ spl72_36 | ~ spl72_103 | ~ spl72_190 | ~ spl72_195 | spl72_233)), inference(subsumption_resolution, [], [f4928, f1326])).
fof(f4928, plain, (~ (e22 = op2(e21, e22)) | (~ spl72_36 | ~ spl72_190 | ~ spl72_195 | spl72_233)), inference(forward_demodulation, [], [f4927, f2069])).
fof(f4927, plain, (~ (op2(e21, e22) = h4(e13)) | (~ spl72_36 | ~ spl72_190 | ~ spl72_195 | spl72_233)), inference(forward_demodulation, [], [f4926, f1009])).
fof(f4926, plain, (~ (op2(e21, e22) = h4(op1(e11, e13))) | (~ spl72_190 | ~ spl72_195 | spl72_233)), inference(forward_demodulation, [], [f4925, f2095])).
fof(f4925, plain, (~ (h4(op1(e11, e13)) = op2(h4(e11), e22)) | (~ spl72_190 | spl72_233)), inference(forward_demodulation, [], [f2301, f2069])).
fof(f4924, plain, (~ spl72_58 | ~ spl72_122 | ~ spl72_195 | ~ spl72_200 | spl72_230), inference(avatar_contradiction_clause, [], [f4923])).
fof(f4923, plain, ($false | (~ spl72_58 | ~ spl72_122 | ~ spl72_195 | ~ spl72_200 | spl72_230)), inference(subsumption_resolution, [], [f4922, f1407])).
fof(f4922, plain, (~ (e21 = op2(e20, e21)) | (~ spl72_58 | ~ spl72_195 | ~ spl72_200 | spl72_230)), inference(forward_demodulation, [], [f4921, f2095])).
fof(f4921, plain, (~ (op2(e20, e21) = h4(e11)) | (~ spl72_58 | ~ spl72_195 | ~ spl72_200 | spl72_230)), inference(forward_demodulation, [], [f4920, f1103])).
fof(f4920, plain, (~ (op2(e20, e21) = h4(op1(e10, e11))) | (~ spl72_195 | ~ spl72_200 | spl72_230)), inference(forward_demodulation, [], [f4919, f2120])).
fof(f4919, plain, (~ (h4(op1(e10, e11)) = op2(h4(e10), e21)) | (~ spl72_195 | spl72_230)), inference(forward_demodulation, [], [f2289, f2095])).
fof(f4914, plain, (~ spl72_51 | ~ spl72_120 | ~ spl72_190 | ~ spl72_200 | spl72_231), inference(avatar_contradiction_clause, [], [f4913])).
fof(f4913, plain, ($false | (~ spl72_51 | ~ spl72_120 | ~ spl72_190 | ~ spl72_200 | spl72_231)), inference(subsumption_resolution, [], [f4911, f1398])).
fof(f1398, plain, ((e23 = op2(e20, e22)) | ~ spl72_120), inference(avatar_component_clause, [], [f1396])).
fof(f1396, plain, (spl72_120 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_120])])).
fof(f4911, plain, (~ (e23 = op2(e20, e22)) | (~ spl72_51 | ~ spl72_190 | ~ spl72_200 | spl72_231)), inference(backward_demodulation, [], [f4906, f2069])).
fof(f4906, plain, (~ (e23 = op2(e20, h4(e13))) | (~ spl72_51 | ~ spl72_200 | spl72_231)), inference(forward_demodulation, [], [f4905, f758])).
fof(f4905, plain, (~ (h4(e12) = op2(e20, h4(e13))) | (~ spl72_51 | ~ spl72_200 | spl72_231)), inference(forward_demodulation, [], [f4904, f1073])).
fof(f4904, plain, (~ (h4(op1(e10, e13)) = op2(e20, h4(e13))) | (~ spl72_200 | spl72_231)), inference(forward_demodulation, [], [f2293, f2120])).
fof(f4895, plain, (spl72_190 | ~ spl72_75 | ~ spl72_195), inference(avatar_split_clause, [], [f4894, f2094, f1205, f2068])).
fof(f4894, plain, ((e22 = h4(e13)) | (~ spl72_75 | ~ spl72_195)), inference(forward_demodulation, [], [f4382, f1207])).
fof(f4382, plain, ((op2(e23, e21) = h4(e13)) | ~ spl72_195), inference(forward_demodulation, [], [f2061, f2095])).
fof(f4893, plain, (~ spl72_190 | ~ spl72_64 | ~ spl72_127 | ~ spl72_200 | spl72_229), inference(avatar_split_clause, [], [f4892, f2283, f2119, f1426, f1126, f2068])).
fof(f1426, plain, (spl72_127 <=> (op2(e20, e20) = e22)), introduced(avatar_definition, [new_symbols(naming, [spl72_127])])).
fof(f4892, plain, (~ (e22 = h4(e13)) | (~ spl72_64 | ~ spl72_127 | ~ spl72_200 | spl72_229)), inference(forward_demodulation, [], [f4891, f1428])).
fof(f1428, plain, ((op2(e20, e20) = e22) | ~ spl72_127), inference(avatar_component_clause, [], [f1426])).
fof(f4891, plain, (~ (op2(e20, e20) = h4(e13)) | (~ spl72_64 | ~ spl72_200 | spl72_229)), inference(forward_demodulation, [], [f4794, f2120])).
fof(f4794, plain, (~ (h4(e13) = op2(h4(e10), h4(e10))) | (~ spl72_64 | spl72_229)), inference(backward_demodulation, [], [f2285, f1128])).
fof(f4886, plain, (~ spl72_226 | ~ spl72_292), inference(avatar_contradiction_clause, [], [f4885])).
fof(f4885, plain, ($false | (~ spl72_226 | ~ spl72_292)), inference(subsumption_resolution, [], [f4884, f368])).
fof(f4884, plain, ((e20 = e23) | (~ spl72_226 | ~ spl72_292)), inference(backward_demodulation, [], [f2676, f2257])).
fof(f2257, plain, ((e20 = h1(e13)) | ~ spl72_226), inference(avatar_component_clause, [], [f2256])).
fof(f4830, plain, (~ spl72_19 | ~ spl72_31), inference(avatar_split_clause, [], [f4825, f986, f935])).
fof(f4825, plain, (~ (e12 = op1(e12, e13)) | ~ spl72_31), inference(backward_demodulation, [], [f302, f988])).
fof(f4813, plain, (~ spl72_49 | ~ spl72_53), inference(avatar_split_clause, [], [f4806, f1080, f1063])).
fof(f4806, plain, (~ (e10 = op1(e10, e13)) | ~ spl72_53), inference(backward_demodulation, [], [f293, f1082])).
fof(f293, plain, ~ (op1(e10, e12) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f4812, plain, (~ spl72_5 | ~ spl72_53), inference(avatar_split_clause, [], [f4805, f1080, f876])).
fof(f4805, plain, (~ (e10 = op1(e13, e12)) | ~ spl72_53), inference(backward_demodulation, [], [f278, f1082])).
fof(f4798, plain, (~ spl72_56 | ~ spl72_64), inference(avatar_split_clause, [], [f4792, f1126, f1092])).
fof(f4792, plain, (~ (e13 = op1(e10, e12)) | ~ spl72_64), inference(backward_demodulation, [], [f289, f1128])).
fof(f289, plain, ~ (op1(e10, e10) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f4797, plain, (~ spl72_16 | ~ spl72_64), inference(avatar_split_clause, [], [f4791, f1126, f922])).
fof(f4791, plain, (~ (e13 = op1(e13, e10)) | ~ spl72_64), inference(backward_demodulation, [], [f266, f1128])).
fof(f4772, plain, (spl72_200 | ~ spl72_105 | ~ spl72_195), inference(avatar_split_clause, [], [f4771, f2094, f1333, f2119])).
fof(f4771, plain, ((e20 = h4(e10)) | (~ spl72_105 | ~ spl72_195)), inference(forward_demodulation, [], [f4384, f1335])).
fof(f4384, plain, ((op2(e21, e21) = h4(e10)) | ~ spl72_195), inference(forward_demodulation, [], [f2062, f2095])).
fof(f4763, plain, (~ spl72_50 | ~ spl72_58), inference(avatar_split_clause, [], [f4762, f1101, f1067])).
fof(f1067, plain, (spl72_50 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_50])])).
fof(f4762, plain, (~ (e11 = op1(e10, e13)) | ~ spl72_58), inference(forward_demodulation, [], [f292, f1103])).
fof(f292, plain, ~ (op1(e10, e11) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f4730, plain, (~ spl72_46 | ~ spl72_110 | ~ spl72_195 | ~ spl72_200 | spl72_232), inference(avatar_contradiction_clause, [], [f4729])).
fof(f4729, plain, ($false | (~ spl72_46 | ~ spl72_110 | ~ spl72_195 | ~ spl72_200 | spl72_232)), inference(subsumption_resolution, [], [f4723, f2095])).
fof(f4723, plain, (~ (e21 = h4(e11)) | (~ spl72_46 | ~ spl72_110 | ~ spl72_195 | ~ spl72_200 | spl72_232)), inference(backward_demodulation, [], [f4683, f1052])).
fof(f4683, plain, (~ (e21 = h4(op1(e11, e10))) | (~ spl72_110 | ~ spl72_195 | ~ spl72_200 | spl72_232)), inference(forward_demodulation, [], [f4226, f1356])).
fof(f1356, plain, ((e21 = op2(e21, e20)) | ~ spl72_110), inference(avatar_component_clause, [], [f1354])).
fof(f4226, plain, (~ (op2(e21, e20) = h4(op1(e11, e10))) | (~ spl72_195 | ~ spl72_200 | spl72_232)), inference(forward_demodulation, [], [f4225, f2095])).
fof(f4225, plain, (~ (h4(op1(e11, e10)) = op2(h4(e11), e20)) | (~ spl72_200 | spl72_232)), inference(forward_demodulation, [], [f2297, f2120])).
fof(f4697, plain, (~ spl72_49 | ~ spl72_113 | spl72_260), inference(avatar_contradiction_clause, [], [f4696])).
fof(f4696, plain, ($false | (~ spl72_49 | ~ spl72_113 | spl72_260)), inference(subsumption_resolution, [], [f4695, f1369])).
fof(f4695, plain, (~ (e20 = op2(e20, e23)) | (~ spl72_49 | spl72_260)), inference(forward_demodulation, [], [f4694, f2054])).
fof(f4694, plain, (~ (op2(e20, e23) = h3(e10)) | (~ spl72_49 | spl72_260)), inference(forward_demodulation, [], [f2472, f1065])).
fof(f2472, plain, (~ (op2(e20, e23) = h3(op1(e10, e13))) | spl72_260), inference(avatar_component_clause, [], [f2470])).
fof(f2470, plain, (spl72_260 <=> (op2(e20, e23) = h3(op1(e10, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_260])])).
fof(f4681, plain, (~ spl72_35 | ~ spl72_19), inference(avatar_split_clause, [], [f4680, f935, f1003])).
fof(f1003, plain, (spl72_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_35])])).
fof(f4680, plain, (~ (e12 = op1(e11, e13)) | ~ spl72_19), inference(forward_demodulation, [], [f285, f937])).
fof(f285, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f4679, plain, (~ spl72_35 | ~ spl72_39), inference(avatar_split_clause, [], [f4678, f1020, f1003])).
fof(f4678, plain, (~ (e12 = op1(e11, e13)) | ~ spl72_39), inference(forward_demodulation, [], [f299, f1022])).
fof(f299, plain, ~ (op1(e11, e12) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f4674, plain, (~ spl72_48 | ~ spl72_16), inference(avatar_split_clause, [], [f4673, f922, f1058])).
fof(f4673, plain, (~ (e13 = op1(e11, e10)) | ~ spl72_16), inference(forward_demodulation, [], [f268, f924])).
fof(f268, plain, ~ (op1(e11, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f4672, plain, (~ spl72_16 | ~ spl72_80 | spl72_251), inference(avatar_contradiction_clause, [], [f4671])).
fof(f4671, plain, ($false | (~ spl72_16 | ~ spl72_80 | spl72_251)), inference(subsumption_resolution, [], [f4670, f1228])).
fof(f4670, plain, (~ (e23 = op2(e23, e20)) | (~ spl72_16 | spl72_251)), inference(forward_demodulation, [], [f4669, f2044])).
fof(f4669, plain, (~ (op2(e23, e20) = h3(e13)) | (~ spl72_16 | spl72_251)), inference(forward_demodulation, [], [f2436, f924])).
fof(f2436, plain, (~ (op2(e23, e20) = h3(op1(e13, e10))) | spl72_251), inference(avatar_component_clause, [], [f2434])).
fof(f2434, plain, (spl72_251 <=> (op2(e23, e20) = h3(op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_251])])).
fof(f4657, plain, (~ spl72_5 | ~ spl72_69 | spl72_249), inference(avatar_contradiction_clause, [], [f4656])).
fof(f4656, plain, ($false | (~ spl72_5 | ~ spl72_69 | spl72_249)), inference(subsumption_resolution, [], [f4654, f2054])).
fof(f4654, plain, (~ (e20 = h3(e10)) | (~ spl72_5 | ~ spl72_69 | spl72_249)), inference(backward_demodulation, [], [f4625, f878])).
fof(f4625, plain, (~ (e20 = h3(op1(e13, e12))) | (~ spl72_69 | spl72_249)), inference(forward_demodulation, [], [f2428, f1182])).
fof(f2428, plain, (~ (op2(e23, e22) = h3(op1(e13, e12))) | spl72_249), inference(avatar_component_clause, [], [f2426])).
fof(f4646, plain, (~ spl72_22 | ~ spl72_24), inference(avatar_contradiction_clause, [], [f4645])).
fof(f4645, plain, ($false | (~ spl72_22 | ~ spl72_24)), inference(subsumption_resolution, [], [f4644, f364])).
fof(f364, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f7])).
fof(f4644, plain, ((e11 = e13) | (~ spl72_22 | ~ spl72_24)), inference(forward_demodulation, [], [f958, f950])).
fof(f958, plain, ((e13 = op1(e12, e12)) | ~ spl72_24), inference(avatar_component_clause, [], [f956])).
fof(f956, plain, (spl72_24 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_24])])).
fof(f4643, plain, (~ spl72_49 | ~ spl72_52), inference(avatar_contradiction_clause, [], [f4642])).
fof(f4642, plain, ($false | (~ spl72_49 | ~ spl72_52)), inference(subsumption_resolution, [], [f4640, f362])).
fof(f362, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f7])).
fof(f4640, plain, ((e10 = e13) | (~ spl72_49 | ~ spl72_52)), inference(backward_demodulation, [], [f1077, f1065])).
fof(f4636, plain, (~ spl72_56 | ~ spl72_120 | spl72_261), inference(avatar_contradiction_clause, [], [f4635])).
fof(f4635, plain, ($false | (~ spl72_56 | ~ spl72_120 | spl72_261)), inference(subsumption_resolution, [], [f4632, f2044])).
fof(f4632, plain, (~ (e23 = h3(e13)) | (~ spl72_56 | ~ spl72_120 | spl72_261)), inference(backward_demodulation, [], [f4628, f1094])).
fof(f4628, plain, (~ (e23 = h3(op1(e10, e12))) | (~ spl72_120 | spl72_261)), inference(forward_demodulation, [], [f2476, f1398])).
fof(f2476, plain, (~ (op2(e20, e22) = h3(op1(e10, e12))) | spl72_261), inference(avatar_component_clause, [], [f2474])).
fof(f2474, plain, (spl72_261 <=> (op2(e20, e22) = h3(op1(e10, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_261])])).
fof(f4618, plain, (~ spl72_8 | ~ spl72_80 | ~ spl72_292 | spl72_293), inference(avatar_contradiction_clause, [], [f4617])).
fof(f4617, plain, ($false | (~ spl72_8 | ~ spl72_80 | ~ spl72_292 | spl72_293)), inference(subsumption_resolution, [], [f4616, f1228])).
fof(f4616, plain, (~ (e23 = op2(e23, e20)) | (~ spl72_8 | ~ spl72_292 | spl72_293)), inference(forward_demodulation, [], [f4615, f2676])).
fof(f4615, plain, (~ (op2(e23, e20) = h1(e13)) | (~ spl72_8 | ~ spl72_292 | spl72_293)), inference(forward_demodulation, [], [f4614, f890])).
fof(f4614, plain, (~ (op2(e23, e20) = h1(op1(e13, e12))) | (~ spl72_292 | spl72_293)), inference(forward_demodulation, [], [f2681, f2676])).
fof(f2681, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | spl72_293), inference(avatar_component_clause, [], [f2679])).
fof(f2679, plain, (spl72_293 <=> (h1(op1(e13, e12)) = op2(h1(e13), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_293])])).
fof(f4609, plain, (~ spl72_14 | ~ spl72_75 | ~ spl72_219 | ~ spl72_224 | spl72_289 | ~ spl72_292), inference(avatar_contradiction_clause, [], [f4608])).
fof(f4608, plain, ($false | (~ spl72_14 | ~ spl72_75 | ~ spl72_219 | ~ spl72_224 | spl72_289 | ~ spl72_292)), inference(subsumption_resolution, [], [f4607, f1207])).
fof(f4607, plain, (~ (e22 = op2(e23, e21)) | (~ spl72_14 | ~ spl72_219 | ~ spl72_224 | spl72_289 | ~ spl72_292)), inference(forward_demodulation, [], [f4606, f2223])).
fof(f4606, plain, (~ (op2(e23, e21) = h1(e11)) | (~ spl72_14 | ~ spl72_224 | spl72_289 | ~ spl72_292)), inference(forward_demodulation, [], [f4605, f916])).
fof(f916, plain, ((e11 = op1(e13, e10)) | ~ spl72_14), inference(avatar_component_clause, [], [f914])).
fof(f914, plain, (spl72_14 <=> (e11 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_14])])).
fof(f4605, plain, (~ (op2(e23, e21) = h1(op1(e13, e10))) | (~ spl72_224 | spl72_289 | ~ spl72_292)), inference(forward_demodulation, [], [f4604, f2676])).
fof(f4604, plain, (~ (h1(op1(e13, e10)) = op2(h1(e13), e21)) | (~ spl72_224 | spl72_289)), inference(forward_demodulation, [], [f2665, f2248])).
fof(f2665, plain, (~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | spl72_289), inference(avatar_component_clause, [], [f2663])).
fof(f2663, plain, (spl72_289 <=> (h1(op1(e13, e10)) = op2(h1(e13), h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_289])])).
fof(f4603, plain, (~ spl72_48 | ~ spl72_92 | ~ spl72_219 | ~ spl72_224 | spl72_287 | ~ spl72_292), inference(avatar_contradiction_clause, [], [f4602])).
fof(f4602, plain, ($false | (~ spl72_48 | ~ spl72_92 | ~ spl72_219 | ~ spl72_224 | spl72_287 | ~ spl72_292)), inference(subsumption_resolution, [], [f4601, f1279])).
fof(f4601, plain, (~ (e23 = op2(e22, e21)) | (~ spl72_48 | ~ spl72_219 | ~ spl72_224 | spl72_287 | ~ spl72_292)), inference(forward_demodulation, [], [f4600, f2676])).
fof(f4600, plain, (~ (op2(e22, e21) = h1(e13)) | (~ spl72_48 | ~ spl72_219 | ~ spl72_224 | spl72_287)), inference(forward_demodulation, [], [f4599, f1060])).
fof(f4599, plain, (~ (op2(e22, e21) = h1(op1(e11, e10))) | (~ spl72_219 | ~ spl72_224 | spl72_287)), inference(forward_demodulation, [], [f4598, f2223])).
fof(f4598, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), e21)) | (~ spl72_224 | spl72_287)), inference(forward_demodulation, [], [f2657, f2248])).
fof(f2657, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | spl72_287), inference(avatar_component_clause, [], [f2655])).
fof(f2655, plain, (spl72_287 <=> (h1(op1(e11, e10)) = op2(h1(e11), h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_287])])).
fof(f4590, plain, (~ spl72_19 | ~ spl72_113 | ~ spl72_292 | spl72_294), inference(avatar_contradiction_clause, [], [f4589])).
fof(f4589, plain, ($false | (~ spl72_19 | ~ spl72_113 | ~ spl72_292 | spl72_294)), inference(subsumption_resolution, [], [f4588, f1369])).
fof(f4588, plain, (~ (e20 = op2(e20, e23)) | (~ spl72_19 | ~ spl72_292 | spl72_294)), inference(forward_demodulation, [], [f4587, f746])).
fof(f4587, plain, (~ (op2(e20, e23) = h1(e12)) | (~ spl72_19 | ~ spl72_292 | spl72_294)), inference(forward_demodulation, [], [f4586, f937])).
fof(f4586, plain, (~ (op2(e20, e23) = h1(op1(e12, e13))) | (~ spl72_292 | spl72_294)), inference(forward_demodulation, [], [f2685, f2676])).
fof(f2685, plain, (~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | spl72_294), inference(avatar_component_clause, [], [f2683])).
fof(f2683, plain, (spl72_294 <=> (h1(op1(e12, e13)) = op2(e20, h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_294])])).
fof(f4585, plain, (~ spl72_11 | ~ spl72_69 | ~ spl72_219 | spl72_290 | ~ spl72_292), inference(avatar_contradiction_clause, [], [f4584])).
fof(f4584, plain, ($false | (~ spl72_11 | ~ spl72_69 | ~ spl72_219 | spl72_290 | ~ spl72_292)), inference(subsumption_resolution, [], [f4583, f1182])).
fof(f4583, plain, (~ (e20 = op2(e23, e22)) | (~ spl72_11 | ~ spl72_219 | spl72_290 | ~ spl72_292)), inference(forward_demodulation, [], [f4582, f746])).
fof(f4582, plain, (~ (op2(e23, e22) = h1(e12)) | (~ spl72_11 | ~ spl72_219 | spl72_290 | ~ spl72_292)), inference(forward_demodulation, [], [f4581, f903])).
fof(f4581, plain, (~ (op2(e23, e22) = h1(op1(e13, e11))) | (~ spl72_219 | spl72_290 | ~ spl72_292)), inference(forward_demodulation, [], [f4580, f2676])).
fof(f4580, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), e22)) | (~ spl72_219 | spl72_290)), inference(forward_demodulation, [], [f2669, f2223])).
fof(f2669, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | spl72_290), inference(avatar_component_clause, [], [f2667])).
fof(f2667, plain, (spl72_290 <=> (h1(op1(e13, e11)) = op2(h1(e13), h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_290])])).
fof(f4576, plain, (~ spl72_1 | ~ spl72_66 | ~ spl72_224 | spl72_291 | ~ spl72_292), inference(avatar_contradiction_clause, [], [f4575])).
fof(f4575, plain, ($false | (~ spl72_1 | ~ spl72_66 | ~ spl72_224 | spl72_291 | ~ spl72_292)), inference(subsumption_resolution, [], [f4574, f1169])).
fof(f1169, plain, ((e21 = op2(e23, e23)) | ~ spl72_66), inference(avatar_component_clause, [], [f1167])).
fof(f4574, plain, (~ (e21 = op2(e23, e23)) | (~ spl72_1 | ~ spl72_224 | spl72_291 | ~ spl72_292)), inference(forward_demodulation, [], [f4573, f2248])).
fof(f4573, plain, (~ (op2(e23, e23) = h1(e10)) | (~ spl72_1 | spl72_291 | ~ spl72_292)), inference(forward_demodulation, [], [f4572, f861])).
fof(f4572, plain, (~ (op2(e23, e23) = h1(op1(e13, e13))) | (spl72_291 | ~ spl72_292)), inference(forward_demodulation, [], [f2673, f2676])).
fof(f2673, plain, (~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | spl72_291), inference(avatar_component_clause, [], [f2671])).
fof(f2671, plain, (spl72_291 <=> (h1(op1(e13, e13)) = op2(h1(e13), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_291])])).
fof(f4564, plain, (~ spl72_39 | ~ spl72_93 | ~ spl72_219 | spl72_298), inference(avatar_contradiction_clause, [], [f4563])).
fof(f4563, plain, ($false | (~ spl72_39 | ~ spl72_93 | ~ spl72_219 | spl72_298)), inference(subsumption_resolution, [], [f4562, f1284])).
fof(f4562, plain, (~ (e20 = op2(e22, e20)) | (~ spl72_39 | ~ spl72_219 | spl72_298)), inference(forward_demodulation, [], [f4561, f746])).
fof(f4561, plain, (~ (op2(e22, e20) = h1(e12)) | (~ spl72_39 | ~ spl72_219 | spl72_298)), inference(forward_demodulation, [], [f4560, f1022])).
fof(f4560, plain, (~ (op2(e22, e20) = h1(op1(e11, e12))) | (~ spl72_219 | spl72_298)), inference(forward_demodulation, [], [f2701, f2223])).
fof(f2701, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | spl72_298), inference(avatar_component_clause, [], [f2699])).
fof(f4555, plain, (~ spl72_34 | ~ spl72_83 | ~ spl72_219 | spl72_288 | ~ spl72_292), inference(avatar_contradiction_clause, [], [f4554])).
fof(f4554, plain, ($false | (~ spl72_34 | ~ spl72_83 | ~ spl72_219 | spl72_288 | ~ spl72_292)), inference(subsumption_resolution, [], [f4553, f1241])).
fof(f4553, plain, (~ (e22 = op2(e22, e23)) | (~ spl72_34 | ~ spl72_219 | spl72_288 | ~ spl72_292)), inference(forward_demodulation, [], [f4552, f2223])).
fof(f4552, plain, (~ (op2(e22, e23) = h1(e11)) | (~ spl72_34 | ~ spl72_219 | spl72_288 | ~ spl72_292)), inference(forward_demodulation, [], [f4551, f1001])).
fof(f4551, plain, (~ (op2(e22, e23) = h1(op1(e11, e13))) | (~ spl72_219 | spl72_288 | ~ spl72_292)), inference(forward_demodulation, [], [f4550, f2223])).
fof(f4550, plain, (~ (h1(op1(e11, e13)) = op2(h1(e11), e23)) | (spl72_288 | ~ spl72_292)), inference(forward_demodulation, [], [f2661, f2676])).
fof(f2661, plain, (~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | spl72_288), inference(avatar_component_clause, [], [f2659])).
fof(f2659, plain, (spl72_288 <=> (h1(op1(e11, e13)) = op2(h1(e11), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_288])])).
fof(f4549, plain, (~ spl72_52 | ~ spl72_100 | ~ spl72_224 | spl72_286 | ~ spl72_292), inference(avatar_contradiction_clause, [], [f4548])).
fof(f4548, plain, ($false | (~ spl72_52 | ~ spl72_100 | ~ spl72_224 | spl72_286 | ~ spl72_292)), inference(subsumption_resolution, [], [f4546, f1313])).
fof(f4546, plain, (~ (e23 = op2(e21, e23)) | (~ spl72_52 | ~ spl72_224 | spl72_286 | ~ spl72_292)), inference(backward_demodulation, [], [f4535, f2248])).
fof(f4535, plain, (~ (e23 = op2(h1(e10), e23)) | (~ spl72_52 | spl72_286 | ~ spl72_292)), inference(forward_demodulation, [], [f4534, f2676])).
fof(f4534, plain, (~ (h1(e13) = op2(h1(e10), e23)) | (~ spl72_52 | spl72_286 | ~ spl72_292)), inference(forward_demodulation, [], [f4533, f1077])).
fof(f4533, plain, (~ (h1(op1(e10, e13)) = op2(h1(e10), e23)) | (spl72_286 | ~ spl72_292)), inference(forward_demodulation, [], [f2653, f2676])).
fof(f2653, plain, (~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | spl72_286), inference(avatar_component_clause, [], [f2651])).
fof(f2651, plain, (spl72_286 <=> (h1(op1(e10, e13)) = op2(h1(e10), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl72_286])])).
fof(f4522, plain, (spl72_224 | ~ spl72_86 | ~ spl72_219), inference(avatar_split_clause, [], [f4521, f2222, f1252, f2247])).
fof(f4521, plain, ((e21 = h1(e10)) | (~ spl72_86 | ~ spl72_219)), inference(forward_demodulation, [], [f4520, f1254])).
fof(f4520, plain, ((op2(e22, e22) = h1(e10)) | ~ spl72_219), inference(forward_demodulation, [], [f2035, f2223])).
fof(f4508, plain, (~ spl72_58 | ~ spl72_103 | ~ spl72_219 | ~ spl72_224 | spl72_285), inference(avatar_contradiction_clause, [], [f4507])).
fof(f4507, plain, ($false | (~ spl72_58 | ~ spl72_103 | ~ spl72_219 | ~ spl72_224 | spl72_285)), inference(subsumption_resolution, [], [f4506, f1326])).
fof(f4506, plain, (~ (e22 = op2(e21, e22)) | (~ spl72_58 | ~ spl72_219 | ~ spl72_224 | spl72_285)), inference(forward_demodulation, [], [f4505, f2223])).
fof(f4505, plain, (~ (op2(e21, e22) = h1(e11)) | (~ spl72_58 | ~ spl72_219 | ~ spl72_224 | spl72_285)), inference(forward_demodulation, [], [f4504, f1103])).
fof(f4504, plain, (~ (op2(e21, e22) = h1(op1(e10, e11))) | (~ spl72_219 | ~ spl72_224 | spl72_285)), inference(forward_demodulation, [], [f4503, f2248])).
fof(f4503, plain, (~ (h1(op1(e10, e11)) = op2(h1(e10), e22)) | (~ spl72_219 | spl72_285)), inference(forward_demodulation, [], [f2649, f2223])).
fof(f2649, plain, (~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | spl72_285), inference(avatar_component_clause, [], [f2647])).
fof(f2647, plain, (spl72_285 <=> (h1(op1(e10, e11)) = op2(h1(e10), h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_285])])).
fof(f4495, plain, (spl72_292 | ~ spl72_120 | ~ spl72_219), inference(avatar_split_clause, [], [f4453, f2222, f1396, f2675])).
fof(f4453, plain, ((e23 = h1(e13)) | (~ spl72_120 | ~ spl72_219)), inference(forward_demodulation, [], [f4452, f1398])).
fof(f4452, plain, ((op2(e20, e22) = h1(e13)) | ~ spl72_219), inference(forward_demodulation, [], [f2034, f2223])).
fof(f4450, plain, (~ spl72_69 | ~ spl72_71), inference(avatar_contradiction_clause, [], [f4449])).
fof(f4449, plain, ($false | (~ spl72_69 | ~ spl72_71)), inference(subsumption_resolution, [], [f4448, f367])).
fof(f367, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f4448, plain, ((e20 = e22) | (~ spl72_69 | ~ spl72_71)), inference(backward_demodulation, [], [f1190, f1182])).
fof(f1190, plain, ((e22 = op2(e23, e22)) | ~ spl72_71), inference(avatar_component_clause, [], [f1188])).
fof(f1188, plain, (spl72_71 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_71])])).
fof(f4445, plain, (~ spl72_72 | ~ spl72_80), inference(avatar_split_clause, [], [f4444, f1226, f1192])).
fof(f4444, plain, (~ (e23 = op2(e23, e22)) | ~ spl72_80), inference(backward_demodulation, [], [f355, f1228])).
fof(f355, plain, ~ (op2(e23, e20) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f4442, plain, (~ spl72_81 | ~ spl72_93), inference(avatar_split_clause, [], [f4439, f1282, f1231])).
fof(f4439, plain, (~ (e20 = op2(e22, e23)) | ~ spl72_93), inference(backward_demodulation, [], [f350, f1284])).
fof(f4441, plain, (~ spl72_77 | ~ spl72_93), inference(avatar_split_clause, [], [f4438, f1282, f1214])).
fof(f4438, plain, (~ (e20 = op2(e23, e20)) | ~ spl72_93), inference(backward_demodulation, [], [f317, f1284])).
fof(f317, plain, ~ (op2(e22, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f4435, plain, (~ spl72_81 | ~ spl72_113), inference(avatar_split_clause, [], [f4433, f1367, f1231])).
fof(f4433, plain, (~ (e20 = op2(e22, e23)) | ~ spl72_113), inference(backward_demodulation, [], [f331, f1369])).
fof(f331, plain, ~ (op2(e20, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f4430, plain, (~ spl72_72 | ~ spl72_120), inference(avatar_split_clause, [], [f4429, f1396, f1192])).
fof(f4429, plain, (~ (e23 = op2(e23, e22)) | ~ spl72_120), inference(backward_demodulation, [], [f326, f1398])).
fof(f4417, plain, (~ spl72_115 | ~ spl72_219), inference(avatar_split_clause, [], [f4407, f2222, f1375])).
fof(f4407, plain, (~ (e22 = op2(e20, e23)) | ~ spl72_219), inference(backward_demodulation, [], [f2033, f2223])).
fof(f4416, plain, (~ spl72_95 | ~ spl72_219), inference(avatar_split_clause, [], [f4404, f2222, f1290])).
fof(f4404, plain, (~ (e22 = op2(e22, e20)) | ~ spl72_219), inference(backward_demodulation, [], [f2029, f2223])).
fof(f2029, plain, ~ (op2(e22, e20) = h1(e11)), inference(backward_demodulation, [], [f313, f748])).
fof(f313, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f4369, plain, (~ spl72_119 | ~ spl72_103), inference(avatar_split_clause, [], [f4368, f1324, f1392])).
fof(f1392, plain, (spl72_119 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_119])])).
fof(f4368, plain, (~ (e22 = op2(e20, e22)) | ~ spl72_103), inference(forward_demodulation, [], [f324, f1326])).
fof(f324, plain, ~ (op2(e20, e22) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f4363, plain, (~ spl72_116 | ~ spl72_52 | spl72_260), inference(avatar_split_clause, [], [f3997, f2470, f1075, f1379])).
fof(f3997, plain, (~ (e23 = op2(e20, e23)) | (~ spl72_52 | spl72_260)), inference(forward_demodulation, [], [f3996, f2044])).
fof(f3996, plain, (~ (op2(e20, e23) = h3(e13)) | (~ spl72_52 | spl72_260)), inference(forward_demodulation, [], [f2472, f1077])).
fof(f4357, plain, (~ spl72_78 | ~ spl72_14 | ~ spl72_202 | spl72_251), inference(avatar_split_clause, [], [f4223, f2434, f2132, f914, f1218])).
fof(f4223, plain, (~ (e21 = op2(e23, e20)) | (~ spl72_14 | ~ spl72_202 | spl72_251)), inference(forward_demodulation, [], [f4222, f2133])).
fof(f4222, plain, (~ (op2(e23, e20) = h3(e11)) | (~ spl72_14 | spl72_251)), inference(forward_demodulation, [], [f2436, f916])).
fof(f4352, plain, (~ spl72_71 | ~ spl72_75), inference(avatar_split_clause, [], [f3796, f1205, f1188])).
fof(f3796, plain, (~ (e22 = op2(e23, e22)) | ~ spl72_75), inference(backward_demodulation, [], [f357, f1207])).
fof(f357, plain, ~ (op2(e23, e21) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f4351, plain, (~ spl72_39 | ~ spl72_77 | spl72_298 | ~ spl72_301), inference(avatar_contradiction_clause, [], [f4350])).
fof(f4350, plain, ($false | (~ spl72_39 | ~ spl72_77 | spl72_298 | ~ spl72_301)), inference(subsumption_resolution, [], [f4349, f1216])).
fof(f1216, plain, ((e20 = op2(e23, e20)) | ~ spl72_77), inference(avatar_component_clause, [], [f1214])).
fof(f4349, plain, (~ (e20 = op2(e23, e20)) | (~ spl72_39 | spl72_298 | ~ spl72_301)), inference(forward_demodulation, [], [f4348, f746])).
fof(f4348, plain, (~ (op2(e23, e20) = h1(e12)) | (~ spl72_39 | spl72_298 | ~ spl72_301)), inference(forward_demodulation, [], [f4347, f1022])).
fof(f4347, plain, (~ (op2(e23, e20) = h1(op1(e11, e12))) | (spl72_298 | ~ spl72_301)), inference(forward_demodulation, [], [f2701, f2724])).
fof(f4340, plain, (~ spl72_53 | ~ spl72_110 | ~ spl72_224 | spl72_300), inference(avatar_contradiction_clause, [], [f4339])).
fof(f4339, plain, ($false | (~ spl72_53 | ~ spl72_110 | ~ spl72_224 | spl72_300)), inference(subsumption_resolution, [], [f4338, f1356])).
fof(f4338, plain, (~ (e21 = op2(e21, e20)) | (~ spl72_53 | ~ spl72_224 | spl72_300)), inference(forward_demodulation, [], [f4337, f2248])).
fof(f4337, plain, (~ (op2(e21, e20) = h1(e10)) | (~ spl72_53 | ~ spl72_224 | spl72_300)), inference(forward_demodulation, [], [f4336, f1082])).
fof(f4336, plain, (~ (op2(e21, e20) = h1(op1(e10, e12))) | (~ spl72_224 | spl72_300)), inference(forward_demodulation, [], [f2709, f2248])).
fof(f2709, plain, (~ (h1(op1(e10, e12)) = op2(h1(e10), e20)) | spl72_300), inference(avatar_component_clause, [], [f2707])).
fof(f2707, plain, (spl72_300 <=> (h1(op1(e10, e12)) = op2(h1(e10), e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_300])])).
fof(f4335, plain, (~ spl72_29 | ~ spl72_122 | ~ spl72_224 | spl72_297), inference(avatar_contradiction_clause, [], [f4334])).
fof(f4334, plain, ($false | (~ spl72_29 | ~ spl72_122 | ~ spl72_224 | spl72_297)), inference(subsumption_resolution, [], [f4333, f1407])).
fof(f4333, plain, (~ (e21 = op2(e20, e21)) | (~ spl72_29 | ~ spl72_224 | spl72_297)), inference(forward_demodulation, [], [f4332, f2248])).
fof(f4332, plain, (~ (op2(e20, e21) = h1(e10)) | (~ spl72_29 | ~ spl72_224 | spl72_297)), inference(forward_demodulation, [], [f4331, f980])).
fof(f4331, plain, (~ (op2(e20, e21) = h1(op1(e12, e10))) | (~ spl72_224 | spl72_297)), inference(forward_demodulation, [], [f2697, f2248])).
fof(f2697, plain, (~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | spl72_297), inference(avatar_component_clause, [], [f2695])).
fof(f2695, plain, (spl72_297 <=> (h1(op1(e12, e10)) = op2(e20, h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl72_297])])).
fof(f4330, plain, (~ spl72_22 | spl72_295), inference(avatar_contradiction_clause, [], [f4329])).
fof(f4329, plain, ($false | (~ spl72_22 | spl72_295)), inference(trivial_inequality_removal, [], [f4328])).
fof(f4328, plain, (~ (h1(e11) = h1(e11)) | (~ spl72_22 | spl72_295)), inference(forward_demodulation, [], [f2689, f950])).
fof(f2689, plain, (~ (h1(e11) = h1(op1(e12, e12))) | spl72_295), inference(avatar_component_clause, [], [f2687])).
fof(f2687, plain, (spl72_295 <=> (h1(e11) = h1(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_295])])).
fof(f4324, plain, (~ spl72_28 | spl72_296), inference(avatar_contradiction_clause, [], [f4323])).
fof(f4323, plain, ($false | (~ spl72_28 | spl72_296)), inference(trivial_inequality_removal, [], [f4322])).
fof(f4322, plain, (~ (h1(e13) = h1(e13)) | (~ spl72_28 | spl72_296)), inference(forward_demodulation, [], [f2693, f975])).
fof(f2693, plain, (~ (h1(e13) = h1(op1(e12, e11))) | spl72_296), inference(avatar_component_clause, [], [f2691])).
fof(f2691, plain, (spl72_296 <=> (h1(e13) = h1(op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_296])])).
fof(f4321, plain, (~ spl72_8 | ~ spl72_95 | ~ spl72_218 | spl72_293), inference(avatar_contradiction_clause, [], [f4320])).
fof(f4320, plain, ($false | (~ spl72_8 | ~ spl72_95 | ~ spl72_218 | spl72_293)), inference(subsumption_resolution, [], [f4319, f1292])).
fof(f4319, plain, (~ (e22 = op2(e22, e20)) | (~ spl72_8 | ~ spl72_218 | spl72_293)), inference(forward_demodulation, [], [f4318, f2217])).
fof(f2217, plain, ((e22 = h1(e13)) | ~ spl72_218), inference(avatar_component_clause, [], [f2216])).
fof(f2216, plain, (spl72_218 <=> (e22 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_218])])).
fof(f4318, plain, (~ (op2(e22, e20) = h1(e13)) | (~ spl72_8 | ~ spl72_218 | spl72_293)), inference(forward_demodulation, [], [f4317, f890])).
fof(f4317, plain, (~ (op2(e22, e20) = h1(op1(e13, e12))) | (~ spl72_218 | spl72_293)), inference(forward_demodulation, [], [f2681, f2217])).
fof(f4310, plain, (~ spl72_19 | ~ spl72_117 | ~ spl72_218 | spl72_294), inference(avatar_contradiction_clause, [], [f4309])).
fof(f4309, plain, ($false | (~ spl72_19 | ~ spl72_117 | ~ spl72_218 | spl72_294)), inference(subsumption_resolution, [], [f4308, f1386])).
fof(f4308, plain, (~ (e20 = op2(e20, e22)) | (~ spl72_19 | ~ spl72_218 | spl72_294)), inference(forward_demodulation, [], [f4307, f746])).
fof(f4307, plain, (~ (op2(e20, e22) = h1(e12)) | (~ spl72_19 | ~ spl72_218 | spl72_294)), inference(forward_demodulation, [], [f4306, f937])).
fof(f4306, plain, (~ (op2(e20, e22) = h1(op1(e12, e13))) | (~ spl72_218 | spl72_294)), inference(forward_demodulation, [], [f2685, f2217])).
fof(f4301, plain, (~ spl72_1 | ~ spl72_86 | ~ spl72_218 | ~ spl72_224 | spl72_291), inference(avatar_contradiction_clause, [], [f4300])).
fof(f4300, plain, ($false | (~ spl72_1 | ~ spl72_86 | ~ spl72_218 | ~ spl72_224 | spl72_291)), inference(subsumption_resolution, [], [f4299, f1254])).
fof(f4299, plain, (~ (e21 = op2(e22, e22)) | (~ spl72_1 | ~ spl72_218 | ~ spl72_224 | spl72_291)), inference(forward_demodulation, [], [f4298, f2248])).
fof(f4298, plain, (~ (op2(e22, e22) = h1(e10)) | (~ spl72_1 | ~ spl72_218 | spl72_291)), inference(forward_demodulation, [], [f4297, f861])).
fof(f4297, plain, (~ (op2(e22, e22) = h1(op1(e13, e13))) | (~ spl72_218 | spl72_291)), inference(forward_demodulation, [], [f2673, f2217])).
fof(f4292, plain, (~ spl72_11 | ~ spl72_81 | ~ spl72_218 | spl72_290 | ~ spl72_301), inference(avatar_contradiction_clause, [], [f4291])).
fof(f4291, plain, ($false | (~ spl72_11 | ~ spl72_81 | ~ spl72_218 | spl72_290 | ~ spl72_301)), inference(subsumption_resolution, [], [f4290, f1233])).
fof(f4290, plain, (~ (e20 = op2(e22, e23)) | (~ spl72_11 | ~ spl72_218 | spl72_290 | ~ spl72_301)), inference(forward_demodulation, [], [f4289, f746])).
fof(f4289, plain, (~ (op2(e22, e23) = h1(e12)) | (~ spl72_11 | ~ spl72_218 | spl72_290 | ~ spl72_301)), inference(forward_demodulation, [], [f4288, f903])).
fof(f4288, plain, (~ (op2(e22, e23) = h1(op1(e13, e11))) | (~ spl72_218 | spl72_290 | ~ spl72_301)), inference(forward_demodulation, [], [f4287, f2217])).
fof(f4287, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), e23)) | (spl72_290 | ~ spl72_301)), inference(forward_demodulation, [], [f2669, f2724])).
fof(f4282, plain, (~ spl72_34 | ~ spl72_72 | ~ spl72_218 | spl72_288 | ~ spl72_301), inference(avatar_contradiction_clause, [], [f4281])).
fof(f4281, plain, ($false | (~ spl72_34 | ~ spl72_72 | ~ spl72_218 | spl72_288 | ~ spl72_301)), inference(subsumption_resolution, [], [f4280, f1194])).
fof(f4280, plain, (~ (e23 = op2(e23, e22)) | (~ spl72_34 | ~ spl72_218 | spl72_288 | ~ spl72_301)), inference(forward_demodulation, [], [f4279, f2724])).
fof(f4279, plain, (~ (op2(e23, e22) = h1(e11)) | (~ spl72_34 | ~ spl72_218 | spl72_288 | ~ spl72_301)), inference(forward_demodulation, [], [f4278, f1001])).
fof(f4278, plain, (~ (op2(e23, e22) = h1(op1(e11, e13))) | (~ spl72_218 | spl72_288 | ~ spl72_301)), inference(forward_demodulation, [], [f4277, f2724])).
fof(f4277, plain, (~ (h1(op1(e11, e13)) = op2(h1(e11), e22)) | (~ spl72_218 | spl72_288)), inference(forward_demodulation, [], [f2661, f2217])).
fof(f4272, plain, (~ spl72_48 | ~ spl72_75 | ~ spl72_218 | ~ spl72_224 | spl72_287 | ~ spl72_301), inference(avatar_contradiction_clause, [], [f4271])).
fof(f4271, plain, ($false | (~ spl72_48 | ~ spl72_75 | ~ spl72_218 | ~ spl72_224 | spl72_287 | ~ spl72_301)), inference(subsumption_resolution, [], [f4270, f1207])).
fof(f4270, plain, (~ (e22 = op2(e23, e21)) | (~ spl72_48 | ~ spl72_218 | ~ spl72_224 | spl72_287 | ~ spl72_301)), inference(forward_demodulation, [], [f4269, f2217])).
fof(f4269, plain, (~ (op2(e23, e21) = h1(e13)) | (~ spl72_48 | ~ spl72_224 | spl72_287 | ~ spl72_301)), inference(forward_demodulation, [], [f4268, f1060])).
fof(f4268, plain, (~ (op2(e23, e21) = h1(op1(e11, e10))) | (~ spl72_224 | spl72_287 | ~ spl72_301)), inference(forward_demodulation, [], [f4267, f2724])).
fof(f4267, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), e21)) | (~ spl72_224 | spl72_287)), inference(forward_demodulation, [], [f2657, f2248])).
fof(f4262, plain, (~ spl72_52 | ~ spl72_103 | ~ spl72_218 | ~ spl72_224 | spl72_286), inference(avatar_contradiction_clause, [], [f4261])).
fof(f4261, plain, ($false | (~ spl72_52 | ~ spl72_103 | ~ spl72_218 | ~ spl72_224 | spl72_286)), inference(subsumption_resolution, [], [f4260, f1326])).
fof(f4260, plain, (~ (e22 = op2(e21, e22)) | (~ spl72_52 | ~ spl72_218 | ~ spl72_224 | spl72_286)), inference(forward_demodulation, [], [f4259, f2217])).
fof(f4259, plain, (~ (op2(e21, e22) = h1(e13)) | (~ spl72_52 | ~ spl72_218 | ~ spl72_224 | spl72_286)), inference(forward_demodulation, [], [f4258, f1077])).
fof(f4258, plain, (~ (op2(e21, e22) = h1(op1(e10, e13))) | (~ spl72_218 | ~ spl72_224 | spl72_286)), inference(forward_demodulation, [], [f4257, f2248])).
fof(f4257, plain, (~ (h1(op1(e10, e13)) = op2(h1(e10), e22)) | (~ spl72_218 | spl72_286)), inference(forward_demodulation, [], [f2653, f2217])).
fof(f4253, plain, (~ spl72_58 | ~ spl72_100 | ~ spl72_224 | spl72_285 | ~ spl72_301), inference(avatar_contradiction_clause, [], [f4252])).
fof(f4252, plain, ($false | (~ spl72_58 | ~ spl72_100 | ~ spl72_224 | spl72_285 | ~ spl72_301)), inference(subsumption_resolution, [], [f4251, f1313])).
fof(f4251, plain, (~ (e23 = op2(e21, e23)) | (~ spl72_58 | ~ spl72_224 | spl72_285 | ~ spl72_301)), inference(forward_demodulation, [], [f4250, f2724])).
fof(f4250, plain, (~ (op2(e21, e23) = h1(e11)) | (~ spl72_58 | ~ spl72_224 | spl72_285 | ~ spl72_301)), inference(forward_demodulation, [], [f4249, f1103])).
fof(f4249, plain, (~ (op2(e21, e23) = h1(op1(e10, e11))) | (~ spl72_224 | spl72_285 | ~ spl72_301)), inference(forward_demodulation, [], [f4248, f2248])).
fof(f4248, plain, (~ (h1(op1(e10, e11)) = op2(h1(e10), e23)) | (spl72_285 | ~ spl72_301)), inference(forward_demodulation, [], [f2649, f2724])).
fof(f4238, plain, (~ spl72_63 | ~ spl72_105 | ~ spl72_224 | spl72_284), inference(avatar_contradiction_clause, [], [f4237])).
fof(f4237, plain, ($false | (~ spl72_63 | ~ spl72_105 | ~ spl72_224 | spl72_284)), inference(subsumption_resolution, [], [f4236, f1335])).
fof(f4236, plain, (~ (e20 = op2(e21, e21)) | (~ spl72_63 | ~ spl72_224 | spl72_284)), inference(forward_demodulation, [], [f4235, f746])).
fof(f4235, plain, (~ (op2(e21, e21) = h1(e12)) | (~ spl72_63 | ~ spl72_224 | spl72_284)), inference(forward_demodulation, [], [f4234, f1124])).
fof(f4234, plain, (~ (op2(e21, e21) = h1(op1(e10, e10))) | (~ spl72_224 | spl72_284)), inference(forward_demodulation, [], [f2645, f2248])).
fof(f2645, plain, (~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))) | spl72_284), inference(avatar_component_clause, [], [f2643])).
fof(f4205, plain, (~ spl72_41 | spl72_299), inference(avatar_contradiction_clause, [], [f4204])).
fof(f4204, plain, ($false | (~ spl72_41 | spl72_299)), inference(trivial_inequality_removal, [], [f4203])).
fof(f4203, plain, (~ (h1(e10) = h1(e10)) | (~ spl72_41 | spl72_299)), inference(forward_demodulation, [], [f2705, f1031])).
fof(f2705, plain, (~ (h1(e10) = h1(op1(e11, e11))) | spl72_299), inference(avatar_component_clause, [], [f2703])).
fof(f2703, plain, (spl72_299 <=> (h1(e10) = h1(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_299])])).
fof(f4200, plain, (~ spl72_48 | ~ spl72_110 | ~ spl72_202 | ~ spl72_259), inference(avatar_contradiction_clause, [], [f4199])).
fof(f4199, plain, ($false | (~ spl72_48 | ~ spl72_110 | ~ spl72_202 | ~ spl72_259)), inference(subsumption_resolution, [], [f4198, f370])).
fof(f370, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f8])).
fof(f4198, plain, ((e21 = e23) | (~ spl72_48 | ~ spl72_110 | ~ spl72_202 | ~ spl72_259)), inference(backward_demodulation, [], [f2044, f4197])).
fof(f4197, plain, ((e21 = h3(e13)) | (~ spl72_48 | ~ spl72_110 | ~ spl72_202 | ~ spl72_259)), inference(forward_demodulation, [], [f3966, f1060])).
fof(f3966, plain, ((e21 = h3(op1(e11, e10))) | (~ spl72_110 | ~ spl72_202 | ~ spl72_259)), inference(forward_demodulation, [], [f3892, f1356])).
fof(f3892, plain, ((op2(e21, e20) = h3(op1(e11, e10))) | (~ spl72_202 | ~ spl72_259)), inference(forward_demodulation, [], [f2467, f2133])).
fof(f2467, plain, ((h3(op1(e11, e10)) = op2(h3(e11), e20)) | ~ spl72_259), inference(avatar_component_clause, [], [f2466])).
fof(f4195, plain, (~ spl72_47 | ~ spl72_39), inference(avatar_split_clause, [], [f4194, f1020, f1054])).
fof(f1054, plain, (spl72_47 <=> (e12 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_47])])).
fof(f4194, plain, (~ (e12 = op1(e11, e10)) | ~ spl72_39), inference(forward_demodulation, [], [f295, f1022])).
fof(f295, plain, ~ (op1(e11, e10) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f4167, plain, (~ spl72_14 | ~ spl72_92 | ~ spl72_218 | ~ spl72_224 | spl72_289 | ~ spl72_301), inference(avatar_contradiction_clause, [], [f4166])).
fof(f4166, plain, ($false | (~ spl72_14 | ~ spl72_92 | ~ spl72_218 | ~ spl72_224 | spl72_289 | ~ spl72_301)), inference(subsumption_resolution, [], [f4162, f2724])).
fof(f4162, plain, (~ (e23 = h1(e11)) | (~ spl72_14 | ~ spl72_92 | ~ spl72_218 | ~ spl72_224 | spl72_289)), inference(backward_demodulation, [], [f3964, f916])).
fof(f3964, plain, (~ (e23 = h1(op1(e13, e10))) | (~ spl72_92 | ~ spl72_218 | ~ spl72_224 | spl72_289)), inference(forward_demodulation, [], [f3963, f1279])).
fof(f3963, plain, (~ (op2(e22, e21) = h1(op1(e13, e10))) | (~ spl72_218 | ~ spl72_224 | spl72_289)), inference(forward_demodulation, [], [f3962, f2217])).
fof(f3962, plain, (~ (h1(op1(e13, e10)) = op2(h1(e13), e21)) | (~ spl72_224 | spl72_289)), inference(forward_demodulation, [], [f2665, f2248])).
fof(f4115, plain, (~ spl72_63 | ~ spl72_128 | ~ spl72_200 | spl72_229), inference(avatar_contradiction_clause, [], [f4114])).
fof(f4114, plain, ($false | (~ spl72_63 | ~ spl72_128 | ~ spl72_200 | spl72_229)), inference(subsumption_resolution, [], [f4107, f758])).
fof(f4107, plain, (~ (e23 = h4(e12)) | (~ spl72_63 | ~ spl72_128 | ~ spl72_200 | spl72_229)), inference(backward_demodulation, [], [f3993, f1124])).
fof(f3993, plain, (~ (e23 = h4(op1(e10, e10))) | (~ spl72_128 | ~ spl72_200 | spl72_229)), inference(forward_demodulation, [], [f3992, f1432])).
fof(f1432, plain, ((op2(e20, e20) = e23) | ~ spl72_128), inference(avatar_component_clause, [], [f1430])).
fof(f3992, plain, (~ (op2(e20, e20) = h4(op1(e10, e10))) | (~ spl72_200 | spl72_229)), inference(forward_demodulation, [], [f2285, f2120])).
fof(f4110, plain, (~ spl72_55 | ~ spl72_63), inference(avatar_split_clause, [], [f4104, f1122, f1088])).
fof(f1088, plain, (spl72_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_55])])).
fof(f4104, plain, (~ (e12 = op1(e10, e12)) | ~ spl72_63), inference(backward_demodulation, [], [f289, f1124])).
fof(f4095, plain, (~ spl72_63 | ~ spl72_47), inference(avatar_split_clause, [], [f4094, f1054, f1122])).
fof(f4094, plain, (~ (op1(e10, e10) = e12) | ~ spl72_47), inference(forward_demodulation, [], [f264, f1056])).
fof(f1056, plain, ((e12 = op1(e11, e10)) | ~ spl72_47), inference(avatar_component_clause, [], [f1054])).
fof(f264, plain, ~ (op1(e10, e10) = op1(e11, e10)), inference(cnf_transformation, [], [f5])).
fof(f4084, plain, (~ spl72_31 | ~ spl72_47), inference(avatar_split_clause, [], [f4083, f1054, f986])).
fof(f4083, plain, (~ (e12 = op1(e12, e10)) | ~ spl72_47), inference(forward_demodulation, [], [f267, f1056])).
fof(f267, plain, ~ (op1(e11, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f4077, plain, (~ spl72_19 | ~ spl72_3), inference(avatar_split_clause, [], [f4076, f867, f935])).
fof(f867, plain, (spl72_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_3])])).
fof(f4076, plain, (~ (e12 = op1(e12, e13)) | ~ spl72_3), inference(forward_demodulation, [], [f287, f869])).
fof(f869, plain, ((e12 = op1(e13, e13)) | ~ spl72_3), inference(avatar_component_clause, [], [f867])).
fof(f287, plain, ~ (op1(e12, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4074, plain, (~ spl72_11 | ~ spl72_3), inference(avatar_split_clause, [], [f4073, f867, f901])).
fof(f4073, plain, (~ (e12 = op1(e13, e11)) | ~ spl72_3), inference(forward_demodulation, [], [f310, f869])).
fof(f310, plain, ~ (op1(e13, e11) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f4047, plain, (~ spl72_16 | ~ spl72_77 | ~ spl72_251), inference(avatar_contradiction_clause, [], [f4046])).
fof(f4046, plain, ($false | (~ spl72_16 | ~ spl72_77 | ~ spl72_251)), inference(subsumption_resolution, [], [f4044, f368])).
fof(f4044, plain, ((e20 = e23) | (~ spl72_16 | ~ spl72_77 | ~ spl72_251)), inference(backward_demodulation, [], [f2044, f4041])).
fof(f4041, plain, ((e20 = h3(e13)) | (~ spl72_16 | ~ spl72_77 | ~ spl72_251)), inference(backward_demodulation, [], [f3960, f924])).
fof(f3960, plain, ((e20 = h3(op1(e13, e10))) | (~ spl72_77 | ~ spl72_251)), inference(forward_demodulation, [], [f2435, f1216])).
fof(f2435, plain, ((op2(e23, e20) = h3(op1(e13, e10))) | ~ spl72_251), inference(avatar_component_clause, [], [f2434])).
fof(f4029, plain, (~ spl72_41 | ~ spl72_44), inference(avatar_contradiction_clause, [], [f4028])).
fof(f4028, plain, ($false | (~ spl72_41 | ~ spl72_44)), inference(subsumption_resolution, [], [f4027, f362])).
fof(f4027, plain, ((e10 = e13) | (~ spl72_41 | ~ spl72_44)), inference(forward_demodulation, [], [f1043, f1031])).
fof(f1043, plain, ((e13 = op1(e11, e11)) | ~ spl72_44), inference(avatar_component_clause, [], [f1041])).
fof(f1041, plain, (spl72_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_44])])).
fof(f4026, plain, (~ spl72_35 | ~ spl72_47), inference(avatar_split_clause, [], [f4022, f1054, f1003])).
fof(f4022, plain, (~ (e12 = op1(e11, e13)) | ~ spl72_47), inference(backward_demodulation, [], [f296, f1056])).
fof(f4006, plain, (~ spl72_46 | ~ spl72_62), inference(avatar_split_clause, [], [f3999, f1118, f1050])).
fof(f3999, plain, (~ (e11 = op1(e11, e10)) | ~ spl72_62), inference(backward_demodulation, [], [f264, f1120])).
fof(f3995, plain, (~ spl72_64 | ~ spl72_52), inference(avatar_split_clause, [], [f3994, f1075, f1126])).
fof(f3994, plain, (~ (op1(e10, e10) = e13) | ~ spl72_52), inference(forward_demodulation, [], [f290, f1077])).
fof(f3989, plain, (~ spl72_60 | ~ spl72_28), inference(avatar_split_clause, [], [f3988, f973, f1109])).
fof(f1109, plain, (spl72_60 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_60])])).
fof(f3988, plain, (~ (e13 = op1(e10, e11)) | ~ spl72_28), inference(forward_demodulation, [], [f271, f975])).
fof(f271, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f3961, plain, (~ spl72_15 | ~ spl72_31), inference(avatar_split_clause, [], [f3266, f986, f918])).
fof(f3266, plain, (~ (e12 = op1(e13, e10)) | ~ spl72_31), inference(forward_demodulation, [], [f269, f988])).
fof(f3956, plain, (~ spl72_8 | ~ spl72_40), inference(avatar_split_clause, [], [f3955, f1024, f888])).
fof(f1024, plain, (spl72_40 <=> (e13 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_40])])).
fof(f3955, plain, (~ (e13 = op1(e13, e12)) | ~ spl72_40), inference(forward_demodulation, [], [f280, f1026])).
fof(f1026, plain, ((e13 = op1(e11, e12)) | ~ spl72_40), inference(avatar_component_clause, [], [f1024])).
fof(f280, plain, ~ (op1(e11, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f3952, plain, (~ spl72_4 | ~ spl72_52), inference(avatar_split_clause, [], [f3951, f1075, f871])).
fof(f3951, plain, (~ (e13 = op1(e13, e13)) | ~ spl72_52), inference(forward_demodulation, [], [f284, f1077])).
fof(f284, plain, ~ (op1(e10, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f3934, plain, (~ spl72_53 | ~ spl72_55), inference(avatar_contradiction_clause, [], [f3933])).
fof(f3933, plain, ($false | (~ spl72_53 | ~ spl72_55)), inference(subsumption_resolution, [], [f3932, f361])).
fof(f361, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f7])).
fof(f3932, plain, ((e10 = e12) | (~ spl72_53 | ~ spl72_55)), inference(forward_demodulation, [], [f1090, f1082])).
fof(f1090, plain, ((e12 = op1(e10, e12)) | ~ spl72_55), inference(avatar_component_clause, [], [f1088])).
fof(f3921, plain, (~ spl72_22 | spl72_253), inference(avatar_contradiction_clause, [], [f3920])).
fof(f3920, plain, ($false | (~ spl72_22 | spl72_253)), inference(trivial_inequality_removal, [], [f3919])).
fof(f3919, plain, (~ (h3(e11) = h3(e11)) | (~ spl72_22 | spl72_253)), inference(forward_demodulation, [], [f2444, f950])).
fof(f2444, plain, (~ (h3(e11) = h3(op1(e12, e12))) | spl72_253), inference(avatar_component_clause, [], [f2442])).
fof(f2442, plain, (spl72_253 <=> (h3(e11) = h3(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl72_253])])).
fof(f3913, plain, (~ spl72_36 | ~ spl72_100 | ~ spl72_202 | spl72_256), inference(avatar_contradiction_clause, [], [f3912])).
fof(f3912, plain, ($false | (~ spl72_36 | ~ spl72_100 | ~ spl72_202 | spl72_256)), inference(subsumption_resolution, [], [f3911, f1313])).
fof(f3911, plain, (~ (e23 = op2(e21, e23)) | (~ spl72_36 | ~ spl72_202 | spl72_256)), inference(forward_demodulation, [], [f3910, f2044])).
fof(f3910, plain, (~ (op2(e21, e23) = h3(e13)) | (~ spl72_36 | ~ spl72_202 | spl72_256)), inference(forward_demodulation, [], [f3909, f1009])).
fof(f3909, plain, (~ (op2(e21, e23) = h3(op1(e11, e13))) | (~ spl72_202 | spl72_256)), inference(forward_demodulation, [], [f2456, f2133])).
fof(f3908, plain, (~ spl72_17 | ~ spl72_81 | spl72_252), inference(avatar_contradiction_clause, [], [f3907])).
fof(f3907, plain, ($false | (~ spl72_17 | ~ spl72_81 | spl72_252)), inference(subsumption_resolution, [], [f3906, f1233])).
fof(f3906, plain, (~ (e20 = op2(e22, e23)) | (~ spl72_17 | spl72_252)), inference(forward_demodulation, [], [f3905, f2054])).
fof(f3905, plain, (~ (op2(e22, e23) = h3(e10)) | (~ spl72_17 | spl72_252)), inference(forward_demodulation, [], [f2440, f929])).
fof(f3901, plain, (~ spl72_41 | spl72_258), inference(avatar_contradiction_clause, [], [f3900])).
fof(f3900, plain, ($false | (~ spl72_41 | spl72_258)), inference(subsumption_resolution, [], [f3899, f2054])).
fof(f3899, plain, (~ (e20 = h3(e10)) | (~ spl72_41 | spl72_258)), inference(forward_demodulation, [], [f2464, f1031])).
fof(f2464, plain, (~ (e20 = h3(op1(e11, e11))) | spl72_258), inference(avatar_component_clause, [], [f2462])).
fof(f2462, plain, (spl72_258 <=> (e20 = h3(op1(e11, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_258])])).
fof(f3898, plain, (~ spl72_13 | ~ spl72_77 | spl72_251), inference(avatar_contradiction_clause, [], [f3897])).
fof(f3897, plain, ($false | (~ spl72_13 | ~ spl72_77 | spl72_251)), inference(subsumption_resolution, [], [f3896, f1216])).
fof(f3896, plain, (~ (e20 = op2(e23, e20)) | (~ spl72_13 | spl72_251)), inference(forward_demodulation, [], [f3895, f2054])).
fof(f3895, plain, (~ (op2(e23, e20) = h3(e10)) | (~ spl72_13 | spl72_251)), inference(forward_demodulation, [], [f2436, f912])).
fof(f3888, plain, (~ spl72_46 | ~ spl72_110 | ~ spl72_202 | spl72_259), inference(avatar_contradiction_clause, [], [f3887])).
fof(f3887, plain, ($false | (~ spl72_46 | ~ spl72_110 | ~ spl72_202 | spl72_259)), inference(subsumption_resolution, [], [f3886, f1356])).
fof(f3886, plain, (~ (e21 = op2(e21, e20)) | (~ spl72_46 | ~ spl72_202 | spl72_259)), inference(forward_demodulation, [], [f3885, f2133])).
fof(f3885, plain, (~ (op2(e21, e20) = h3(e11)) | (~ spl72_46 | ~ spl72_202 | spl72_259)), inference(forward_demodulation, [], [f3884, f1052])).
fof(f3884, plain, (~ (op2(e21, e20) = h3(op1(e11, e10))) | (~ spl72_202 | spl72_259)), inference(forward_demodulation, [], [f2468, f2133])).
fof(f3883, plain, (~ spl72_11 | ~ spl72_75 | ~ spl72_202 | spl72_250), inference(avatar_contradiction_clause, [], [f3882])).
fof(f3882, plain, ($false | (~ spl72_11 | ~ spl72_75 | ~ spl72_202 | spl72_250)), inference(subsumption_resolution, [], [f3881, f1207])).
fof(f3881, plain, (~ (e22 = op2(e23, e21)) | (~ spl72_11 | ~ spl72_202 | spl72_250)), inference(forward_demodulation, [], [f3880, f754])).
fof(f754, plain, (e22 = h3(e12)), inference(cnf_transformation, [], [f16])).
fof(f3880, plain, (~ (op2(e23, e21) = h3(e12)) | (~ spl72_11 | ~ spl72_202 | spl72_250)), inference(forward_demodulation, [], [f3879, f903])).
fof(f3879, plain, (~ (op2(e23, e21) = h3(op1(e13, e11))) | (~ spl72_202 | spl72_250)), inference(forward_demodulation, [], [f2432, f2133])).
fof(f2432, plain, (~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | spl72_250), inference(avatar_component_clause, [], [f2430])).
fof(f2430, plain, (spl72_250 <=> (h3(op1(e13, e11)) = op2(e23, h3(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_250])])).
fof(f3874, plain, (~ spl72_51 | ~ spl72_115 | spl72_260), inference(avatar_contradiction_clause, [], [f3873])).
fof(f3873, plain, ($false | (~ spl72_51 | ~ spl72_115 | spl72_260)), inference(subsumption_resolution, [], [f3872, f1377])).
fof(f3872, plain, (~ (e22 = op2(e20, e23)) | (~ spl72_51 | spl72_260)), inference(forward_demodulation, [], [f3871, f754])).
fof(f3871, plain, (~ (op2(e20, e23) = h3(e12)) | (~ spl72_51 | spl72_260)), inference(forward_demodulation, [], [f2472, f1073])).
fof(f3870, plain, (~ spl72_8 | ~ spl72_72 | spl72_249), inference(avatar_contradiction_clause, [], [f3869])).
fof(f3869, plain, ($false | (~ spl72_8 | ~ spl72_72 | spl72_249)), inference(subsumption_resolution, [], [f3868, f1194])).
fof(f3868, plain, (~ (e23 = op2(e23, e22)) | (~ spl72_8 | spl72_249)), inference(forward_demodulation, [], [f3867, f2044])).
fof(f3867, plain, (~ (op2(e23, e22) = h3(e13)) | (~ spl72_8 | spl72_249)), inference(forward_demodulation, [], [f2428, f890])).
fof(f3862, plain, (~ spl72_53 | ~ spl72_117 | spl72_261), inference(avatar_contradiction_clause, [], [f3861])).
fof(f3861, plain, ($false | (~ spl72_53 | ~ spl72_117 | spl72_261)), inference(subsumption_resolution, [], [f3860, f1386])).
fof(f3860, plain, (~ (e20 = op2(e20, e22)) | (~ spl72_53 | spl72_261)), inference(forward_demodulation, [], [f3859, f2054])).
fof(f3859, plain, (~ (op2(e20, e22) = h3(e10)) | (~ spl72_53 | spl72_261)), inference(forward_demodulation, [], [f2476, f1082])).
fof(f3858, plain, (~ spl72_2 | ~ spl72_195 | ~ spl72_202 | spl72_248), inference(avatar_contradiction_clause, [], [f3857])).
fof(f3857, plain, ($false | (~ spl72_2 | ~ spl72_195 | ~ spl72_202 | spl72_248)), inference(subsumption_resolution, [], [f3856, f2133])).
fof(f3856, plain, (~ (e21 = h3(e11)) | (~ spl72_2 | ~ spl72_195 | spl72_248)), inference(forward_demodulation, [], [f3855, f2095])).
fof(f3855, plain, (~ (h3(e11) = h4(e11)) | (~ spl72_2 | spl72_248)), inference(forward_demodulation, [], [f2424, f865])).
fof(f3850, plain, (~ spl72_58 | ~ spl72_122 | ~ spl72_202 | spl72_262), inference(avatar_contradiction_clause, [], [f3849])).
fof(f3849, plain, ($false | (~ spl72_58 | ~ spl72_122 | ~ spl72_202 | spl72_262)), inference(subsumption_resolution, [], [f3848, f1407])).
fof(f3848, plain, (~ (e21 = op2(e20, e21)) | (~ spl72_58 | ~ spl72_202 | spl72_262)), inference(forward_demodulation, [], [f3847, f2133])).
fof(f3847, plain, (~ (op2(e20, e21) = h3(e11)) | (~ spl72_58 | ~ spl72_202 | spl72_262)), inference(forward_demodulation, [], [f3846, f1103])).
fof(f3846, plain, (~ (op2(e20, e21) = h3(op1(e10, e11))) | (~ spl72_202 | spl72_262)), inference(forward_demodulation, [], [f2480, f2133])).
fof(f2480, plain, (~ (h3(op1(e10, e11)) = op2(e20, h3(e11))) | spl72_262), inference(avatar_component_clause, [], [f2478])).
fof(f2478, plain, (spl72_262 <=> (h3(op1(e10, e11)) = op2(e20, h3(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_262])])).
fof(f3837, plain, (~ spl72_64 | spl72_263 | ~ spl72_301), inference(avatar_contradiction_clause, [], [f3836])).
fof(f3836, plain, ($false | (~ spl72_64 | spl72_263 | ~ spl72_301)), inference(subsumption_resolution, [], [f3835, f2724])).
fof(f3835, plain, (~ (e23 = h1(e11)) | (~ spl72_64 | spl72_263)), inference(forward_demodulation, [], [f3834, f2044])).
fof(f3834, plain, (~ (h1(e11) = h3(e13)) | (~ spl72_64 | spl72_263)), inference(forward_demodulation, [], [f2484, f1128])).
fof(f3821, plain, (~ spl72_39 | ~ spl72_103 | ~ spl72_202 | spl72_257), inference(avatar_contradiction_clause, [], [f3820])).
fof(f3820, plain, ($false | (~ spl72_39 | ~ spl72_103 | ~ spl72_202 | spl72_257)), inference(subsumption_resolution, [], [f3819, f1326])).
fof(f3819, plain, (~ (e22 = op2(e21, e22)) | (~ spl72_39 | ~ spl72_202 | spl72_257)), inference(forward_demodulation, [], [f3818, f754])).
fof(f3818, plain, (~ (op2(e21, e22) = h3(e12)) | (~ spl72_39 | ~ spl72_202 | spl72_257)), inference(forward_demodulation, [], [f3817, f1022])).
fof(f3817, plain, (~ (op2(e21, e22) = h3(op1(e11, e12))) | (~ spl72_202 | spl72_257)), inference(forward_demodulation, [], [f2460, f2133])).
fof(f2460, plain, (~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | spl72_257), inference(avatar_component_clause, [], [f2458])).
fof(f3803, plain, (spl72_224 | ~ spl72_66 | ~ spl72_301), inference(avatar_split_clause, [], [f3802, f2723, f1167, f2247])).
fof(f3802, plain, ((e21 = h1(e10)) | (~ spl72_66 | ~ spl72_301)), inference(backward_demodulation, [], [f3651, f1169])).
fof(f3651, plain, ((op2(e23, e23) = h1(e10)) | ~ spl72_301), inference(backward_demodulation, [], [f2035, f2724])).
fof(f3792, plain, (~ spl72_69 | ~ spl72_77), inference(avatar_split_clause, [], [f3787, f1214, f1180])).
fof(f3787, plain, (~ (e20 = op2(e23, e22)) | ~ spl72_77), inference(backward_demodulation, [], [f355, f1216])).
fof(f3781, plain, (~ spl72_85 | ~ spl72_86), inference(avatar_contradiction_clause, [], [f3780])).
fof(f3780, plain, ($false | (~ spl72_85 | ~ spl72_86)), inference(subsumption_resolution, [], [f3779, f366])).
fof(f3779, plain, ((e20 = e21) | (~ spl72_85 | ~ spl72_86)), inference(backward_demodulation, [], [f1254, f1250])).
fof(f1250, plain, ((e20 = op2(e22, e22)) | ~ spl72_85), inference(avatar_component_clause, [], [f1248])).
fof(f1248, plain, (spl72_85 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_85])])).
fof(f3770, plain, (spl72_218 | ~ spl72_115 | ~ spl72_301), inference(avatar_split_clause, [], [f3767, f2723, f1375, f2216])).
fof(f3767, plain, ((e22 = h1(e13)) | (~ spl72_115 | ~ spl72_301)), inference(backward_demodulation, [], [f3650, f1377])).
fof(f3650, plain, ((op2(e20, e23) = h1(e13)) | ~ spl72_301), inference(backward_demodulation, [], [f2034, f2724])).
fof(f3744, plain, (~ spl72_28 | spl72_254), inference(avatar_contradiction_clause, [], [f3743])).
fof(f3743, plain, ($false | (~ spl72_28 | spl72_254)), inference(subsumption_resolution, [], [f3742, f2044])).
fof(f3742, plain, (~ (e23 = h3(e13)) | (~ spl72_28 | spl72_254)), inference(forward_demodulation, [], [f2448, f975])).
fof(f2448, plain, (~ (e23 = h3(op1(e12, e11))) | spl72_254), inference(avatar_component_clause, [], [f2446])).
fof(f2446, plain, (spl72_254 <=> (e23 = h3(op1(e12, e11)))), introduced(avatar_definition, [new_symbols(naming, [spl72_254])])).
fof(f3719, plain, (~ spl72_96 | ~ spl72_92), inference(avatar_split_clause, [], [f3321, f1277, f1294])).
fof(f1294, plain, (spl72_96 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_96])])).
fof(f3321, plain, (~ (e23 = op2(e22, e20)) | ~ spl72_92), inference(forward_demodulation, [], [f348, f1279])).
fof(f348, plain, ~ (op2(e22, e20) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f3711, plain, (~ spl72_76 | ~ spl72_92), inference(avatar_split_clause, [], [f3710, f1277, f1209])).
fof(f1209, plain, (spl72_76 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_76])])).
fof(f3710, plain, (~ (e23 = op2(e23, e21)) | ~ spl72_92), inference(forward_demodulation, [], [f323, f1279])).
fof(f323, plain, ~ (op2(e22, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f3705, plain, (~ spl72_7 | ~ spl72_11), inference(avatar_split_clause, [], [f3703, f901, f884])).
fof(f884, plain, (spl72_7 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_7])])).
fof(f3703, plain, (~ (e12 = op1(e13, e12)) | ~ spl72_11), inference(backward_demodulation, [], [f309, f903])).
fof(f309, plain, ~ (op1(e13, e11) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f3688, plain, (~ spl72_86 | ~ spl72_88), inference(avatar_contradiction_clause, [], [f3687])).
fof(f3687, plain, ($false | (~ spl72_86 | ~ spl72_88)), inference(subsumption_resolution, [], [f3686, f370])).
fof(f3686, plain, ((e21 = e23) | (~ spl72_86 | ~ spl72_88)), inference(forward_demodulation, [], [f1262, f1254])).
fof(f1262, plain, ((e23 = op2(e22, e22)) | ~ spl72_88), inference(avatar_component_clause, [], [f1260])).
fof(f1260, plain, (spl72_88 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_88])])).
fof(f3685, plain, (~ spl72_98 | ~ spl72_100), inference(avatar_contradiction_clause, [], [f3684])).
fof(f3684, plain, ($false | (~ spl72_98 | ~ spl72_100)), inference(subsumption_resolution, [], [f3683, f370])).
fof(f3683, plain, ((e21 = e23) | (~ spl72_98 | ~ spl72_100)), inference(forward_demodulation, [], [f1313, f1305])).
fof(f3681, plain, (~ spl72_105 | ~ spl72_107), inference(avatar_contradiction_clause, [], [f3680])).
fof(f3680, plain, ($false | (~ spl72_105 | ~ spl72_107)), inference(subsumption_resolution, [], [f3679, f367])).
fof(f3679, plain, ((e20 = e22) | (~ spl72_105 | ~ spl72_107)), inference(forward_demodulation, [], [f1343, f1335])).
fof(f1343, plain, ((e22 = op2(e21, e21)) | ~ spl72_107), inference(avatar_component_clause, [], [f1341])).
fof(f1341, plain, (spl72_107 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_107])])).
fof(f3645, plain, (spl72_191 | ~ spl72_67), inference(avatar_split_clause, [], [f3487, f1171, f2074])).
fof(f1171, plain, (spl72_67 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_67])])).
fof(f3487, plain, ((e22 = h4(e11)) | ~ spl72_67), inference(backward_demodulation, [], [f760, f1173])).
fof(f1173, plain, ((e22 = op2(e23, e23)) | ~ spl72_67), inference(avatar_component_clause, [], [f1171])).
fof(f3619, plain, (~ spl72_73 | ~ spl72_215), inference(avatar_split_clause, [], [f3503, f2202, f1197])).
fof(f1197, plain, (spl72_73 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_73])])).
fof(f3503, plain, (~ (e20 = op2(e23, e21)) | ~ spl72_215), inference(forward_demodulation, [], [f2038, f2203])).
fof(f2038, plain, ~ (op2(e23, e21) = h2(e11)), inference(backward_demodulation, [], [f322, f752])).
fof(f322, plain, ~ (op2(e21, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f3603, plain, (~ spl72_3 | ~ spl72_7), inference(avatar_split_clause, [], [f3602, f884, f867])).
fof(f3602, plain, (~ (e12 = op1(e13, e13)) | ~ spl72_7), inference(backward_demodulation, [], [f311, f886])).
fof(f886, plain, ((e12 = op1(e13, e12)) | ~ spl72_7), inference(avatar_component_clause, [], [f884])).
fof(f3580, plain, (~ spl72_41 | ~ spl72_43), inference(avatar_contradiction_clause, [], [f3579])).
fof(f3579, plain, ($false | (~ spl72_41 | ~ spl72_43)), inference(subsumption_resolution, [], [f3578, f361])).
fof(f3578, plain, ((e10 = e12) | (~ spl72_41 | ~ spl72_43)), inference(forward_demodulation, [], [f1039, f1031])).
fof(f1039, plain, ((e12 = op1(e11, e11)) | ~ spl72_43), inference(avatar_component_clause, [], [f1037])).
fof(f1037, plain, (spl72_43 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_43])])).
fof(f3575, plain, (~ spl72_35 | ~ spl72_51), inference(avatar_split_clause, [], [f3572, f1071, f1003])).
fof(f3572, plain, (~ (e12 = op1(e11, e13)) | ~ spl72_51), inference(backward_demodulation, [], [f282, f1073])).
fof(f3569, plain, (~ spl72_53 | ~ spl72_56), inference(avatar_contradiction_clause, [], [f3568])).
fof(f3568, plain, ($false | (~ spl72_53 | ~ spl72_56)), inference(subsumption_resolution, [], [f3566, f362])).
fof(f3566, plain, ((e10 = e13) | (~ spl72_53 | ~ spl72_56)), inference(backward_demodulation, [], [f1094, f1082])).
fof(f3553, plain, (~ spl72_78 | ~ spl72_80), inference(avatar_contradiction_clause, [], [f3552])).
fof(f3552, plain, ($false | (~ spl72_78 | ~ spl72_80)), inference(subsumption_resolution, [], [f3551, f370])).
fof(f3551, plain, ((e21 = e23) | (~ spl72_78 | ~ spl72_80)), inference(forward_demodulation, [], [f1228, f1220])).
fof(f3548, plain, (~ spl72_91 | ~ spl72_92), inference(avatar_contradiction_clause, [], [f3547])).
fof(f3547, plain, ($false | (~ spl72_91 | ~ spl72_92)), inference(subsumption_resolution, [], [f3546, f371])).
fof(f371, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f3546, plain, ((e22 = e23) | (~ spl72_91 | ~ spl72_92)), inference(backward_demodulation, [], [f1279, f1275])).
fof(f1275, plain, ((e22 = op2(e22, e21)) | ~ spl72_91), inference(avatar_component_clause, [], [f1273])).
fof(f1273, plain, (spl72_91 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_91])])).
fof(f3537, plain, (spl72_273 | ~ spl72_112 | ~ spl72_215), inference(avatar_split_clause, [], [f3535, f2202, f1362, f2565])).
fof(f3535, plain, ((e23 = h2(e13)) | (~ spl72_112 | ~ spl72_215)), inference(backward_demodulation, [], [f3341, f1364])).
fof(f3520, plain, (spl72_219 | ~ spl72_127), inference(avatar_split_clause, [], [f3519, f1426, f2222])).
fof(f3519, plain, ((e22 = h1(e11)) | ~ spl72_127), inference(backward_demodulation, [], [f748, f1428])).
fof(f3517, plain, (~ spl72_122 | ~ spl72_114), inference(avatar_split_clause, [], [f3516, f1371, f1405])).
fof(f1371, plain, (spl72_114 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_114])])).
fof(f3516, plain, (~ (e21 = op2(e20, e21)) | ~ spl72_114), inference(forward_demodulation, [], [f340, f1373])).
fof(f1373, plain, ((e21 = op2(e20, e23)) | ~ spl72_114), inference(avatar_component_clause, [], [f1371])).
fof(f340, plain, ~ (op2(e20, e21) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f3496, plain, (~ spl72_70 | ~ spl72_202), inference(avatar_split_clause, [], [f3495, f2132, f1184])).
fof(f1184, plain, (spl72_70 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_70])])).
fof(f3495, plain, (~ (e21 = op2(e23, e22)) | ~ spl72_202), inference(forward_demodulation, [], [f2047, f2133])).
fof(f2047, plain, ~ (op2(e23, e22) = h3(e11)), inference(backward_demodulation, [], [f329, f756])).
fof(f329, plain, ~ (op2(e22, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f3476, plain, (~ spl72_38 | ~ spl72_22), inference(avatar_split_clause, [], [f3289, f948, f1016])).
fof(f1016, plain, (spl72_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_38])])).
fof(f3289, plain, (~ (e11 = op1(e11, e12)) | ~ spl72_22), inference(forward_demodulation, [], [f279, f950])).
fof(f279, plain, ~ (op1(e11, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f3472, plain, (~ spl72_20 | ~ spl72_28), inference(avatar_split_clause, [], [f3271, f973, f939])).
fof(f939, plain, (spl72_20 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_20])])).
fof(f3271, plain, (~ (e13 = op1(e12, e13)) | ~ spl72_28), inference(backward_demodulation, [], [f304, f975])).
fof(f304, plain, ~ (op1(e12, e11) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f3470, plain, (~ spl72_12 | ~ spl72_28), inference(avatar_split_clause, [], [f3272, f973, f905])).
fof(f905, plain, (spl72_12 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_12])])).
fof(f3272, plain, (~ (e13 = op1(e13, e11)) | ~ spl72_28), inference(backward_demodulation, [], [f275, f975])).
fof(f275, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f3468, plain, (~ spl72_3 | ~ spl72_4), inference(avatar_contradiction_clause, [], [f3467])).
fof(f3467, plain, ($false | (~ spl72_3 | ~ spl72_4)), inference(subsumption_resolution, [], [f3466, f365])).
fof(f365, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f7])).
fof(f3466, plain, ((e12 = e13) | (~ spl72_3 | ~ spl72_4)), inference(backward_demodulation, [], [f873, f869])).
fof(f873, plain, ((e13 = op1(e13, e13)) | ~ spl72_4), inference(avatar_component_clause, [], [f871])).
fof(f3448, plain, (~ spl72_21 | ~ spl72_22), inference(avatar_contradiction_clause, [], [f3447])).
fof(f3447, plain, ($false | (~ spl72_21 | ~ spl72_22)), inference(subsumption_resolution, [], [f3446, f360])).
fof(f360, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f7])).
fof(f3446, plain, ((e10 = e11) | (~ spl72_21 | ~ spl72_22)), inference(backward_demodulation, [], [f950, f946])).
fof(f946, plain, ((e10 = op1(e12, e12)) | ~ spl72_21), inference(avatar_component_clause, [], [f944])).
fof(f944, plain, (spl72_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_21])])).
fof(f3438, plain, (~ spl72_46 | ~ spl72_48), inference(avatar_contradiction_clause, [], [f3437])).
fof(f3437, plain, ($false | (~ spl72_46 | ~ spl72_48)), inference(subsumption_resolution, [], [f3436, f364])).
fof(f3436, plain, ((e11 = e13) | (~ spl72_46 | ~ spl72_48)), inference(forward_demodulation, [], [f1060, f1052])).
fof(f3435, plain, (~ spl72_49 | ~ spl72_50), inference(avatar_contradiction_clause, [], [f3434])).
fof(f3434, plain, ($false | (~ spl72_49 | ~ spl72_50)), inference(subsumption_resolution, [], [f3433, f360])).
fof(f3433, plain, ((e10 = e11) | (~ spl72_49 | ~ spl72_50)), inference(backward_demodulation, [], [f1069, f1065])).
fof(f1069, plain, ((e11 = op1(e10, e13)) | ~ spl72_50), inference(avatar_component_clause, [], [f1067])).
fof(f3426, plain, (~ spl72_58 | ~ spl72_59), inference(avatar_contradiction_clause, [], [f3425])).
fof(f3425, plain, ($false | (~ spl72_58 | ~ spl72_59)), inference(subsumption_resolution, [], [f3424, f363])).
fof(f3424, plain, ((e11 = e12) | (~ spl72_58 | ~ spl72_59)), inference(forward_demodulation, [], [f1107, f1103])).
fof(f1107, plain, ((e12 = op1(e10, e11)) | ~ spl72_59), inference(avatar_component_clause, [], [f1105])).
fof(f3400, plain, (~ spl72_93 | ~ spl72_95), inference(avatar_contradiction_clause, [], [f3399])).
fof(f3399, plain, ($false | (~ spl72_93 | ~ spl72_95)), inference(subsumption_resolution, [], [f3398, f367])).
fof(f3398, plain, ((e20 = e22) | (~ spl72_93 | ~ spl72_95)), inference(backward_demodulation, [], [f1292, f1284])).
fof(f3381, plain, (spl72_210 | ~ spl72_110 | ~ spl72_215), inference(avatar_split_clause, [], [f3379, f2202, f1354, f2175])).
fof(f3379, plain, ((e21 = h2(e13)) | (~ spl72_110 | ~ spl72_215)), inference(backward_demodulation, [], [f3341, f1356])).
fof(f3369, plain, (~ spl72_117 | ~ spl72_119), inference(avatar_contradiction_clause, [], [f3368])).
fof(f3368, plain, ($false | (~ spl72_117 | ~ spl72_119)), inference(subsumption_resolution, [], [f3367, f367])).
fof(f3367, plain, ((e20 = e22) | (~ spl72_117 | ~ spl72_119)), inference(backward_demodulation, [], [f1394, f1386])).
fof(f1394, plain, ((e22 = op2(e20, e22)) | ~ spl72_119), inference(avatar_component_clause, [], [f1392])).
fof(f3364, plain, (~ spl72_219 | ~ spl72_119), inference(avatar_split_clause, [], [f3362, f1392, f2222])).
fof(f3362, plain, (~ (e22 = h1(e11)) | ~ spl72_119), inference(backward_demodulation, [], [f2032, f1394])).
fof(f2032, plain, ~ (op2(e20, e22) = h1(e11)), inference(backward_demodulation, [], [f337, f748])).
fof(f337, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f3354, plain, (spl72_301 | ~ spl72_128), inference(avatar_split_clause, [], [f3353, f1430, f2723])).
fof(f3353, plain, ((e23 = h1(e11)) | ~ spl72_128), inference(backward_demodulation, [], [f748, f1432])).
fof(f3348, plain, (~ spl72_97 | ~ spl72_215), inference(avatar_split_clause, [], [f3340, f2202, f1299])).
fof(f1299, plain, (spl72_97 <=> (e20 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_97])])).
fof(f3340, plain, (~ (e20 = op2(e21, e23)) | ~ spl72_215), inference(backward_demodulation, [], [f2041, f2203])).
fof(f2041, plain, ~ (op2(e21, e23) = h2(e11)), inference(backward_demodulation, [], [f346, f752])).
fof(f346, plain, ~ (op2(e21, e21) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f3347, plain, (~ spl72_109 | ~ spl72_215), inference(avatar_split_clause, [], [f3339, f2202, f1350])).
fof(f1350, plain, (spl72_109 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_109])])).
fof(f3339, plain, (~ (e20 = op2(e21, e20)) | ~ spl72_215), inference(backward_demodulation, [], [f2039, f2203])).
fof(f2039, plain, ~ (op2(e21, e20) = h2(e11)), inference(backward_demodulation, [], [f342, f752])).
fof(f342, plain, ~ (op2(e21, e20) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f3346, plain, (~ spl72_121 | ~ spl72_215), inference(avatar_split_clause, [], [f3338, f2202, f1401])).
fof(f1401, plain, (spl72_121 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_121])])).
fof(f3338, plain, (~ (e20 = op2(e20, e21)) | ~ spl72_215), inference(backward_demodulation, [], [f2036, f2203])).
fof(f2036, plain, ~ (op2(e20, e21) = h2(e11)), inference(backward_demodulation, [], [f318, f752])).
fof(f318, plain, ~ (op2(e20, e21) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f3336, plain, (~ spl72_124 | ~ spl72_92), inference(avatar_split_clause, [], [f3335, f1277, f1413])).
fof(f3335, plain, (~ (e23 = op2(e20, e21)) | ~ spl72_92), inference(forward_demodulation, [], [f319, f1279])).
fof(f319, plain, ~ (op2(e20, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f3330, plain, (~ spl72_111 | ~ spl72_95), inference(avatar_split_clause, [], [f3329, f1290, f1358])).
fof(f1358, plain, (spl72_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_111])])).
fof(f3329, plain, (~ (e22 = op2(e21, e20)) | ~ spl72_95), inference(forward_demodulation, [], [f315, f1292])).
fof(f315, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f3328, plain, (~ spl72_111 | ~ spl72_103), inference(avatar_split_clause, [], [f3327, f1324, f1358])).
fof(f3327, plain, (~ (e22 = op2(e21, e20)) | ~ spl72_103), inference(forward_demodulation, [], [f343, f1326])).
fof(f343, plain, ~ (op2(e21, e20) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f3324, plain, (~ spl72_99 | ~ spl72_103), inference(avatar_split_clause, [], [f3323, f1324, f1307])).
fof(f1307, plain, (spl72_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_99])])).
fof(f3323, plain, (~ (e22 = op2(e21, e23)) | ~ spl72_103), inference(forward_demodulation, [], [f347, f1326])).
fof(f347, plain, ~ (op2(e21, e22) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f3318, plain, (~ spl72_31 | ~ spl72_95 | spl72_255), inference(avatar_contradiction_clause, [], [f3317])).
fof(f3317, plain, ($false | (~ spl72_31 | ~ spl72_95 | spl72_255)), inference(subsumption_resolution, [], [f3316, f1292])).
fof(f3316, plain, (~ (e22 = op2(e22, e20)) | (~ spl72_31 | spl72_255)), inference(forward_demodulation, [], [f3315, f754])).
fof(f3315, plain, (~ (op2(e22, e20) = h3(e12)) | (~ spl72_31 | spl72_255)), inference(forward_demodulation, [], [f2452, f988])).
fof(f3314, plain, (~ spl72_82 | ~ spl72_202), inference(avatar_split_clause, [], [f3313, f2132, f1235])).
fof(f1235, plain, (spl72_82 <=> (e21 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_82])])).
fof(f3313, plain, (~ (e21 = op2(e22, e23)) | ~ spl72_202), inference(forward_demodulation, [], [f2050, f2133])).
fof(f2050, plain, ~ (op2(e22, e23) = h3(e11)), inference(backward_demodulation, [], [f353, f756])).
fof(f353, plain, ~ (op2(e22, e22) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f3312, plain, (~ spl72_84 | ~ spl72_92), inference(avatar_split_clause, [], [f3311, f1277, f1243])).
fof(f1243, plain, (spl72_84 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_84])])).
fof(f3311, plain, (~ (e23 = op2(e22, e23)) | ~ spl72_92), inference(forward_demodulation, [], [f352, f1279])).
fof(f352, plain, ~ (op2(e22, e21) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f3306, plain, (spl72_195 | ~ spl72_66), inference(avatar_split_clause, [], [f3305, f1167, f2094])).
fof(f3305, plain, ((e21 = h4(e11)) | ~ spl72_66), inference(forward_demodulation, [], [f760, f1169])).
fof(f3279, plain, (~ spl72_55 | ~ spl72_39), inference(avatar_split_clause, [], [f3278, f1020, f1088])).
fof(f3278, plain, (~ (e12 = op1(e10, e12)) | ~ spl72_39), inference(forward_demodulation, [], [f276, f1022])).
fof(f276, plain, ~ (op1(e10, e12) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f3277, plain, (~ spl72_39 | ~ spl72_40), inference(avatar_contradiction_clause, [], [f3276])).
fof(f3276, plain, ($false | (~ spl72_39 | ~ spl72_40)), inference(subsumption_resolution, [], [f3275, f365])).
fof(f3275, plain, ((e12 = e13) | (~ spl72_39 | ~ spl72_40)), inference(forward_demodulation, [], [f1026, f1022])).
fof(f3265, plain, (~ spl72_14 | ~ spl72_46), inference(avatar_split_clause, [], [f3264, f1050, f914])).
fof(f3264, plain, (~ (e11 = op1(e13, e10)) | ~ spl72_46), inference(forward_demodulation, [], [f268, f1052])).
fof(f3260, plain, (~ spl72_2 | ~ spl72_4), inference(avatar_contradiction_clause, [], [f3259])).
fof(f3259, plain, ($false | (~ spl72_2 | ~ spl72_4)), inference(subsumption_resolution, [], [f3258, f364])).
fof(f3258, plain, ((e11 = e13) | (~ spl72_2 | ~ spl72_4)), inference(backward_demodulation, [], [f873, f865])).
fof(f3245, plain, (~ spl72_27 | ~ spl72_28), inference(avatar_contradiction_clause, [], [f3244])).
fof(f3244, plain, ($false | (~ spl72_27 | ~ spl72_28)), inference(subsumption_resolution, [], [f3243, f365])).
fof(f3243, plain, ((e12 = e13) | (~ spl72_27 | ~ spl72_28)), inference(backward_demodulation, [], [f975, f971])).
fof(f971, plain, ((e12 = op1(e12, e11)) | ~ spl72_27), inference(avatar_component_clause, [], [f969])).
fof(f969, plain, (spl72_27 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_27])])).
fof(f3239, plain, (~ spl72_28 | ~ spl72_32), inference(avatar_split_clause, [], [f3236, f990, f973])).
fof(f990, plain, (spl72_32 <=> (e13 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_32])])).
fof(f3236, plain, (~ (e13 = op1(e12, e11)) | ~ spl72_32), inference(backward_demodulation, [], [f300, f992])).
fof(f992, plain, ((e13 = op1(e12, e10)) | ~ spl72_32), inference(avatar_component_clause, [], [f990])).
fof(f300, plain, ~ (op1(e12, e10) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f3214, plain, (~ spl72_65 | ~ spl72_66), inference(avatar_contradiction_clause, [], [f3213])).
fof(f3213, plain, ($false | (~ spl72_65 | ~ spl72_66)), inference(subsumption_resolution, [], [f3212, f366])).
fof(f3212, plain, ((e20 = e21) | (~ spl72_65 | ~ spl72_66)), inference(forward_demodulation, [], [f1169, f1165])).
fof(f3211, plain, (~ spl72_65 | ~ spl72_67), inference(avatar_contradiction_clause, [], [f3210])).
fof(f3210, plain, ($false | (~ spl72_65 | ~ spl72_67)), inference(subsumption_resolution, [], [f3209, f367])).
fof(f3209, plain, ((e20 = e22) | (~ spl72_65 | ~ spl72_67)), inference(forward_demodulation, [], [f1173, f1165])).
fof(f3207, plain, (~ spl72_69 | ~ spl72_72), inference(avatar_contradiction_clause, [], [f3206])).
fof(f3206, plain, ($false | (~ spl72_69 | ~ spl72_72)), inference(subsumption_resolution, [], [f3205, f368])).
fof(f3205, plain, ((e20 = e23) | (~ spl72_69 | ~ spl72_72)), inference(backward_demodulation, [], [f1194, f1182])).
fof(f3186, plain, (~ spl72_103 | ~ spl72_104), inference(avatar_contradiction_clause, [], [f3185])).
fof(f3185, plain, ($false | (~ spl72_103 | ~ spl72_104)), inference(subsumption_resolution, [], [f3184, f371])).
fof(f3184, plain, ((e22 = e23) | (~ spl72_103 | ~ spl72_104)), inference(backward_demodulation, [], [f1330, f1326])).
fof(f1330, plain, ((e23 = op2(e21, e22)) | ~ spl72_104), inference(avatar_component_clause, [], [f1328])).
fof(f1328, plain, (spl72_104 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_104])])).
fof(f3181, plain, (spl72_211 | ~ spl72_106), inference(avatar_split_clause, [], [f3180, f1337, f2182])).
fof(f1337, plain, (spl72_106 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_106])])).
fof(f3180, plain, ((e21 = h2(e11)) | ~ spl72_106), inference(backward_demodulation, [], [f752, f1339])).
fof(f1339, plain, ((e21 = op2(e21, e21)) | ~ spl72_106), inference(avatar_component_clause, [], [f1337])).
fof(f3177, plain, (~ spl72_104 | ~ spl72_120), inference(avatar_split_clause, [], [f3175, f1396, f1328])).
fof(f3175, plain, (~ (e23 = op2(e21, e22)) | ~ spl72_120), inference(backward_demodulation, [], [f324, f1398])).
fof(f3173, plain, (spl72_223 | ~ spl72_126), inference(avatar_split_clause, [], [f3172, f1422, f2242])).
fof(f1422, plain, (spl72_126 <=> (op2(e20, e20) = e21)), introduced(avatar_definition, [new_symbols(naming, [spl72_126])])).
fof(f3172, plain, ((e21 = h1(e11)) | ~ spl72_126), inference(backward_demodulation, [], [f748, f1424])).
fof(f1424, plain, ((op2(e20, e20) = e21) | ~ spl72_126), inference(avatar_component_clause, [], [f1422])).
fof(f3169, plain, (spl72_92 | ~ spl72_202), inference(avatar_split_clause, [], [f3162, f2132, f1277])).
fof(f3162, plain, ((e23 = op2(e22, e21)) | ~ spl72_202), inference(backward_demodulation, [], [f2052, f2133])).
fof(f2052, plain, (e23 = op2(e22, h3(e11))), inference(backward_demodulation, [], [f745, f756])).
fof(f3168, plain, (spl72_105 | ~ spl72_202), inference(avatar_split_clause, [], [f3161, f2132, f1333])).
fof(f3161, plain, ((e20 = op2(e21, e21)) | ~ spl72_202), inference(backward_demodulation, [], [f2051, f2133])).
fof(f3166, plain, (~ spl72_94 | ~ spl72_202), inference(avatar_split_clause, [], [f3159, f2132, f1286])).
fof(f1286, plain, (spl72_94 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_94])])).
fof(f3159, plain, (~ (e21 = op2(e22, e20)) | ~ spl72_202), inference(backward_demodulation, [], [f2048, f2133])).
fof(f2048, plain, ~ (op2(e22, e20) = h3(e11)), inference(backward_demodulation, [], [f349, f756])).
fof(f349, plain, ~ (op2(e22, e20) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f3165, plain, (~ spl72_102 | ~ spl72_202), inference(avatar_split_clause, [], [f3158, f2132, f1320])).
fof(f1320, plain, (spl72_102 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_102])])).
fof(f3158, plain, (~ (e21 = op2(e21, e22)) | ~ spl72_202), inference(backward_demodulation, [], [f2046, f2133])).
fof(f2046, plain, ~ (op2(e21, e22) = h3(e11)), inference(backward_demodulation, [], [f327, f756])).
fof(f327, plain, ~ (op2(e21, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f3164, plain, (~ spl72_118 | ~ spl72_202), inference(avatar_split_clause, [], [f3157, f2132, f1388])).
fof(f1388, plain, (spl72_118 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_118])])).
fof(f3157, plain, (~ (e21 = op2(e20, e22)) | ~ spl72_202), inference(backward_demodulation, [], [f2045, f2133])).
fof(f2045, plain, ~ (op2(e20, e22) = h3(e11)), inference(backward_demodulation, [], [f325, f756])).
fof(f325, plain, ~ (op2(e20, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f3143, plain, (~ spl72_104 | ~ spl72_72), inference(avatar_split_clause, [], [f3142, f1192, f1328])).
fof(f3142, plain, (~ (e23 = op2(e21, e22)) | ~ spl72_72), inference(forward_demodulation, [], [f328, f1194])).
fof(f328, plain, ~ (op2(e21, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f3116, plain, (~ spl72_191 | ~ spl72_75), inference(avatar_split_clause, [], [f3115, f1205, f2074])).
fof(f3115, plain, (~ (e22 = h4(e11)) | ~ spl72_75), inference(forward_demodulation, [], [f2059, f1207])).
fof(f2059, plain, ~ (op2(e23, e21) = h4(e11)), inference(backward_demodulation, [], [f358, f760])).
fof(f358, plain, ~ (op2(e23, e21) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f3110, plain, (spl72_199 | ~ spl72_65), inference(avatar_split_clause, [], [f3109, f1163, f2114])).
fof(f3109, plain, ((e20 = h4(e11)) | ~ spl72_65), inference(forward_demodulation, [], [f760, f1165])).
fof(f3102, plain, (~ spl72_54 | ~ spl72_22), inference(avatar_split_clause, [], [f3101, f948, f1084])).
fof(f1084, plain, (spl72_54 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_54])])).
fof(f3101, plain, (~ (e11 = op1(e10, e12)) | ~ spl72_22), inference(forward_demodulation, [], [f277, f950])).
fof(f277, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f3093, plain, (~ spl72_49 | ~ spl72_17), inference(avatar_split_clause, [], [f3092, f927, f1063])).
fof(f3092, plain, (~ (e10 = op1(e10, e13)) | ~ spl72_17), inference(forward_demodulation, [], [f283, f929])).
fof(f283, plain, ~ (op1(e10, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f3081, plain, (~ spl72_41 | ~ spl72_37), inference(avatar_split_clause, [], [f3080, f1012, f1029])).
fof(f1012, plain, (spl72_37 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_37])])).
fof(f3080, plain, (~ (e10 = op1(e11, e11)) | ~ spl72_37), inference(forward_demodulation, [], [f297, f1014])).
fof(f1014, plain, ((e10 = op1(e11, e12)) | ~ spl72_37), inference(avatar_component_clause, [], [f1012])).
fof(f297, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f3077, plain, (spl72_41 | ~ spl72_22), inference(avatar_split_clause, [], [f2993, f948, f1029])).
fof(f2993, plain, ((e10 = op1(e11, e11)) | ~ spl72_22), inference(backward_demodulation, [], [f740, f950])).
fof(f740, plain, (e10 = op1(op1(e12, e12), op1(e12, e12))), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e13 = op1(e12, op1(e12, e12))) & (e11 = op1(e12, e12)) & (e10 = op1(op1(e12, e12), op1(e12, e12)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax12)).
fof(f3070, plain, (~ spl72_30 | ~ spl72_22), inference(avatar_split_clause, [], [f3069, f948, f982])).
fof(f982, plain, (spl72_30 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_30])])).
fof(f3069, plain, (~ (e11 = op1(e12, e10)) | ~ spl72_22), inference(forward_demodulation, [], [f301, f950])).
fof(f301, plain, ~ (op1(e12, e10) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f3060, plain, (spl72_28 | ~ spl72_22), inference(avatar_split_clause, [], [f2994, f948, f973])).
fof(f2994, plain, ((e13 = op1(e12, e11)) | ~ spl72_22), inference(backward_demodulation, [], [f742, f950])).
fof(f742, plain, (e13 = op1(e12, op1(e12, e12))), inference(cnf_transformation, [], [f12])).
fof(f3038, plain, (~ spl72_7 | ~ spl72_8), inference(avatar_contradiction_clause, [], [f3037])).
fof(f3037, plain, ($false | (~ spl72_7 | ~ spl72_8)), inference(subsumption_resolution, [], [f3036, f365])).
fof(f3036, plain, ((e12 = e13) | (~ spl72_7 | ~ spl72_8)), inference(backward_demodulation, [], [f890, f886])).
fof(f3033, plain, (~ spl72_10 | ~ spl72_11), inference(avatar_contradiction_clause, [], [f3032])).
fof(f3032, plain, ($false | (~ spl72_10 | ~ spl72_11)), inference(subsumption_resolution, [], [f3031, f363])).
fof(f3031, plain, ((e11 = e12) | (~ spl72_10 | ~ spl72_11)), inference(backward_demodulation, [], [f903, f899])).
fof(f899, plain, ((e11 = op1(e13, e11)) | ~ spl72_10), inference(avatar_component_clause, [], [f897])).
fof(f3016, plain, (~ spl72_4 | ~ spl72_16), inference(avatar_split_clause, [], [f3012, f922, f871])).
fof(f3012, plain, (~ (e13 = op1(e13, e13)) | ~ spl72_16), inference(backward_demodulation, [], [f308, f924])).
fof(f308, plain, ~ (op1(e13, e10) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f3009, plain, (~ spl72_17 | ~ spl72_19), inference(avatar_contradiction_clause, [], [f3008])).
fof(f3008, plain, ($false | (~ spl72_17 | ~ spl72_19)), inference(subsumption_resolution, [], [f3007, f361])).
fof(f3007, plain, ((e10 = e12) | (~ spl72_17 | ~ spl72_19)), inference(backward_demodulation, [], [f937, f929])).
fof(f2998, plain, (~ spl72_18 | ~ spl72_22), inference(avatar_split_clause, [], [f2992, f948, f931])).
fof(f931, plain, (spl72_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_18])])).
fof(f2992, plain, (~ (e11 = op1(e12, e13)) | ~ spl72_22), inference(backward_demodulation, [], [f305, f950])).
fof(f305, plain, ~ (op1(e12, e12) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f2997, plain, (~ spl72_6 | ~ spl72_22), inference(avatar_split_clause, [], [f2991, f948, f880])).
fof(f880, plain, (spl72_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl72_6])])).
fof(f2991, plain, (~ (e11 = op1(e13, e12)) | ~ spl72_22), inference(backward_demodulation, [], [f281, f950])).
fof(f281, plain, ~ (op1(e12, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f2966, plain, (~ spl72_4 | ~ spl72_36), inference(avatar_split_clause, [], [f2964, f1007, f871])).
fof(f2964, plain, (~ (e13 = op1(e13, e13)) | ~ spl72_36), inference(backward_demodulation, [], [f286, f1009])).
fof(f2956, plain, (~ spl72_33 | ~ spl72_41), inference(avatar_split_clause, [], [f2949, f1029, f995])).
fof(f995, plain, (spl72_33 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl72_33])])).
fof(f2949, plain, (~ (e10 = op1(e11, e13)) | ~ spl72_41), inference(backward_demodulation, [], [f298, f1031])).
fof(f298, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f2943, plain, (~ spl72_41 | ~ spl72_45), inference(avatar_split_clause, [], [f2938, f1046, f1029])).
fof(f1046, plain, (spl72_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_45])])).
fof(f2938, plain, (~ (e10 = op1(e11, e11)) | ~ spl72_45), inference(backward_demodulation, [], [f294, f1048])).
fof(f1048, plain, ((e10 = op1(e11, e10)) | ~ spl72_45), inference(avatar_component_clause, [], [f1046])).
fof(f294, plain, ~ (op1(e11, e10) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f2931, plain, (~ spl72_51 | ~ spl72_52), inference(avatar_contradiction_clause, [], [f2930])).
fof(f2930, plain, ($false | (~ spl72_51 | ~ spl72_52)), inference(subsumption_resolution, [], [f2929, f365])).
fof(f2929, plain, ((e12 = e13) | (~ spl72_51 | ~ spl72_52)), inference(backward_demodulation, [], [f1077, f1073])).
fof(f2910, plain, (~ spl72_41 | ~ spl72_57), inference(avatar_split_clause, [], [f2905, f1097, f1029])).
fof(f1097, plain, (spl72_57 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_57])])).
fof(f2905, plain, (~ (e10 = op1(e11, e11)) | ~ spl72_57), inference(backward_demodulation, [], [f270, f1099])).
fof(f1099, plain, ((e10 = op1(e10, e11)) | ~ spl72_57), inference(avatar_component_clause, [], [f1097])).
fof(f270, plain, ~ (op1(e10, e11) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f2904, plain, (~ spl72_49 | ~ spl72_61), inference(avatar_split_clause, [], [f2894, f1114, f1063])).
fof(f1114, plain, (spl72_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl72_61])])).
fof(f2894, plain, (~ (e10 = op1(e10, e13)) | ~ spl72_61), inference(backward_demodulation, [], [f290, f1116])).
fof(f1116, plain, ((e10 = op1(e10, e10)) | ~ spl72_61), inference(avatar_component_clause, [], [f1114])).
fof(f2903, plain, (~ spl72_53 | ~ spl72_61), inference(avatar_split_clause, [], [f2893, f1114, f1080])).
fof(f2893, plain, (~ (e10 = op1(e10, e12)) | ~ spl72_61), inference(backward_demodulation, [], [f289, f1116])).
fof(f2900, plain, (~ spl72_29 | ~ spl72_61), inference(avatar_split_clause, [], [f2890, f1114, f978])).
fof(f2890, plain, (~ (e10 = op1(e12, e10)) | ~ spl72_61), inference(backward_demodulation, [], [f265, f1116])).
fof(f2888, plain, (~ spl72_65 | ~ spl72_68), inference(avatar_contradiction_clause, [], [f2887])).
fof(f2887, plain, ($false | (~ spl72_65 | ~ spl72_68)), inference(subsumption_resolution, [], [f2886, f368])).
fof(f2886, plain, ((e20 = e23) | (~ spl72_65 | ~ spl72_68)), inference(backward_demodulation, [], [f1177, f1165])).
fof(f1177, plain, ((e23 = op2(e23, e23)) | ~ spl72_68), inference(avatar_component_clause, [], [f1175])).
fof(f1175, plain, (spl72_68 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl72_68])])).
fof(f2883, plain, (spl72_246 | ~ spl72_68), inference(avatar_split_clause, [], [f2882, f1175, f2374])).
fof(f2374, plain, (spl72_246 <=> (e23 = h4(e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_246])])).
fof(f2882, plain, ((e23 = h4(e11)) | ~ spl72_68), inference(backward_demodulation, [], [f760, f1177])).
fof(f2881, plain, (~ spl72_71 | ~ spl72_72), inference(avatar_contradiction_clause, [], [f2880])).
fof(f2880, plain, ($false | (~ spl72_71 | ~ spl72_72)), inference(subsumption_resolution, [], [f2879, f371])).
fof(f2879, plain, ((e22 = e23) | (~ spl72_71 | ~ spl72_72)), inference(backward_demodulation, [], [f1194, f1190])).
fof(f2877, plain, (~ spl72_246 | ~ spl72_72), inference(avatar_split_clause, [], [f2875, f1192, f2374])).
fof(f2875, plain, (~ (e23 = h4(e11)) | ~ spl72_72), inference(backward_demodulation, [], [f2060, f1194])).
fof(f2060, plain, ~ (op2(e23, e22) = h4(e11)), inference(backward_demodulation, [], [f359, f760])).
fof(f359, plain, ~ (op2(e23, e22) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f2874, plain, (~ spl72_74 | ~ spl72_75), inference(avatar_contradiction_clause, [], [f2873])).
fof(f2873, plain, ($false | (~ spl72_74 | ~ spl72_75)), inference(subsumption_resolution, [], [f2872, f369])).
fof(f2872, plain, ((e21 = e22) | (~ spl72_74 | ~ spl72_75)), inference(backward_demodulation, [], [f1207, f1203])).
fof(f1203, plain, ((e21 = op2(e23, e21)) | ~ spl72_74), inference(avatar_component_clause, [], [f1201])).
fof(f2854, plain, (~ spl72_301 | ~ spl72_80), inference(avatar_split_clause, [], [f2850, f1226, f2723])).
fof(f2850, plain, (~ (e23 = h1(e11)) | ~ spl72_80), inference(backward_demodulation, [], [f2030, f1228])).
fof(f2853, plain, (~ spl72_246 | ~ spl72_80), inference(avatar_split_clause, [], [f2849, f1226, f2374])).
fof(f2849, plain, (~ (e23 = h4(e11)) | ~ spl72_80), inference(backward_demodulation, [], [f2058, f1228])).
fof(f2841, plain, (spl72_202 | ~ spl72_86), inference(avatar_split_clause, [], [f2840, f1252, f2132])).
fof(f2840, plain, ((e21 = h3(e11)) | ~ spl72_86), inference(backward_demodulation, [], [f756, f1254])).
fof(f2835, plain, (~ spl72_215 | ~ spl72_89), inference(avatar_split_clause, [], [f2831, f1265, f2202])).
fof(f1265, plain, (spl72_89 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl72_89])])).
fof(f2831, plain, (~ (e20 = h2(e11)) | ~ spl72_89), inference(backward_demodulation, [], [f2037, f1267])).
fof(f1267, plain, ((e20 = op2(e22, e21)) | ~ spl72_89), inference(avatar_component_clause, [], [f1265])).
fof(f2037, plain, ~ (op2(e22, e21) = h2(e11)), inference(backward_demodulation, [], [f321, f752])).
fof(f321, plain, ~ (op2(e21, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f2827, plain, (~ spl72_227 | ~ spl72_93), inference(avatar_split_clause, [], [f2821, f1282, f2263])).
fof(f2821, plain, (~ (e20 = h1(e11)) | ~ spl72_93), inference(backward_demodulation, [], [f2029, f1284])).
fof(f2803, plain, (~ spl72_215 | ~ spl72_101), inference(avatar_split_clause, [], [f2799, f1316, f2202])).
fof(f1316, plain, (spl72_101 <=> (e20 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl72_101])])).
fof(f2799, plain, (~ (e20 = h2(e11)) | ~ spl72_101), inference(backward_demodulation, [], [f2040, f1318])).
fof(f1318, plain, ((e20 = op2(e21, e22)) | ~ spl72_101), inference(avatar_component_clause, [], [f1316])).
fof(f2040, plain, ~ (op2(e21, e22) = h2(e11)), inference(backward_demodulation, [], [f345, f752])).
fof(f345, plain, ~ (op2(e21, e21) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f2797, plain, (spl72_215 | ~ spl72_105), inference(avatar_split_clause, [], [f2796, f1333, f2202])).
fof(f2796, plain, ((e20 = h2(e11)) | ~ spl72_105), inference(backward_demodulation, [], [f752, f1335])).
fof(f2765, plain, (~ spl72_227 | ~ spl72_117), inference(avatar_split_clause, [], [f2760, f1384, f2263])).
fof(f2760, plain, (~ (e20 = h1(e11)) | ~ spl72_117), inference(backward_demodulation, [], [f2032, f1386])).
fof(f2746, plain, (spl72_227 | ~ spl72_125), inference(avatar_split_clause, [], [f2745, f1418, f2263])).
fof(f1418, plain, (spl72_125 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl72_125])])).
fof(f2745, plain, ((e20 = h1(e11)) | ~ spl72_125), inference(backward_demodulation, [], [f748, f1420])).
fof(f1420, plain, ((e20 = op2(e20, e20)) | ~ spl72_125), inference(avatar_component_clause, [], [f1418])).
fof(f2726, plain, (~ spl72_284 | ~ spl72_285 | ~ spl72_286 | ~ spl72_287 | ~ spl72_288 | ~ spl72_289 | ~ spl72_290 | ~ spl72_291 | spl72_225 | spl72_221 | spl72_217 | ~ spl72_301 | ~ spl72_293 | ~ spl72_294 | ~ spl72_295 | ~ spl72_296 | ~ spl72_297 | ~ spl72_298 | ~ spl72_299 | ~ spl72_300), inference(avatar_split_clause, [], [f2721, f2707, f2703, f2699, f2695, f2691, f2687, f2683, f2679, f2723, f2212, f2232, f2252, f2671, f2667, f2663, f2659, f2655, f2651, f2647, f2643])).
fof(f2252, plain, (spl72_225 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl72_225])])).
fof(f2232, plain, (spl72_221 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl72_221])])).
fof(f2212, plain, (spl72_217 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl72_217])])).
fof(f2721, plain, (~ (h1(op1(e10, e12)) = op2(h1(e10), e20)) | ~ (h1(e10) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e12, e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2720, f746])).
fof(f2720, plain, (~ (h1(e10) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e12, e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2719, f2035])).
fof(f2719, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e12, e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2718, f746])).
fof(f2718, plain, (~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e12, e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2717, f746])).
fof(f2717, plain, (~ (h1(e13) = h1(op1(e12, e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2716, f2034])).
fof(f2716, plain, (~ (h1(op1(e12, e11)) = op2(e20, h1(e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2715, f746])).
fof(f2715, plain, (~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2714, f748])).
fof(f2714, plain, (~ (op2(e20, e20) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2713, f746])).
fof(f2713, plain, (~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2712, f746])).
fof(f2712, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f811, f746])).
fof(f811, plain, (~ (e23 = h1(e11)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(cnf_transformation, [], [f95])).
fof(f95, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP68 | sP67 | sP66 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f20, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83])).
fof(f83, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP60), inference(usedef, [], [e83])).
fof(e83, plain, (sP60 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f84, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP61), inference(usedef, [], [e84])).
fof(e84, plain, (sP61 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f85, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP62), inference(usedef, [], [e85])).
fof(e85, plain, (sP62 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f86, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP63), inference(usedef, [], [e86])).
fof(e86, plain, (sP63 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f87, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP64), inference(usedef, [], [e87])).
fof(e87, plain, (sP64 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f88, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP65), inference(usedef, [], [e88])).
fof(e88, plain, (sP65 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f89, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP66), inference(usedef, [], [e89])).
fof(e89, plain, (sP66 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f90, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP67), inference(usedef, [], [e90])).
fof(e90, plain, (sP67 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f91, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP68), inference(usedef, [], [e91])).
fof(e91, plain, (sP68 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f92, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP69), inference(usedef, [], [e92])).
fof(e92, plain, (sP69 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f93, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP70), inference(usedef, [], [e93])).
fof(e93, plain, (sP70 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f94, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP71), inference(usedef, [], [e94])).
fof(e94, plain, (sP71 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f20, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f18])).
fof(f18, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', co1)).
fof(f2710, plain, (~ spl72_284 | ~ spl72_285 | ~ spl72_286 | ~ spl72_287 | ~ spl72_288 | ~ spl72_289 | ~ spl72_290 | ~ spl72_291 | spl72_225 | spl72_221 | spl72_217 | ~ spl72_292 | ~ spl72_293 | ~ spl72_294 | ~ spl72_295 | ~ spl72_296 | ~ spl72_297 | ~ spl72_298 | ~ spl72_299 | ~ spl72_300), inference(avatar_split_clause, [], [f2641, f2707, f2703, f2699, f2695, f2691, f2687, f2683, f2679, f2675, f2212, f2232, f2252, f2671, f2667, f2663, f2659, f2655, f2651, f2647, f2643])).
fof(f2641, plain, (~ (h1(op1(e10, e12)) = op2(h1(e10), e20)) | ~ (h1(e10) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e12, e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2640, f746])).
fof(f2640, plain, (~ (h1(e10) = h1(op1(e11, e11))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e12, e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2639, f2035])).
fof(f2639, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e12, e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2638, f746])).
fof(f2638, plain, (~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(e13) = h1(op1(e12, e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2637, f746])).
fof(f2637, plain, (~ (h1(e13) = h1(op1(e12, e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2636, f2034])).
fof(f2636, plain, (~ (h1(op1(e12, e11)) = op2(e20, h1(e11))) | ~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2635, f746])).
fof(f2635, plain, (~ (h1(e11) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2634, f748])).
fof(f2634, plain, (~ (op2(e20, e20) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2633, f746])).
fof(f2633, plain, (~ (h1(op1(e12, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2632, f746])).
fof(f2632, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f813, f746])).
fof(f813, plain, (~ (e23 = h1(e13)) | sP62 | sP61 | sP60 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(cnf_transformation, [], [f95])).
fof(f2600, plain, (~ spl72_265 | ~ spl72_266 | ~ spl72_267 | ~ spl72_268 | ~ spl72_269 | ~ spl72_270 | ~ spl72_271 | ~ spl72_272 | spl72_213 | spl72_209 | spl72_205 | ~ spl72_273 | ~ spl72_274 | ~ spl72_275 | ~ spl72_276 | ~ spl72_277 | ~ spl72_278 | ~ spl72_279 | ~ spl72_280 | ~ spl72_281), inference(avatar_split_clause, [], [f2531, f2597, f2593, f2589, f2585, f2581, f2577, f2573, f2569, f2565, f2151, f2171, f2192, f2561, f2557, f2553, f2549, f2545, f2541, f2537, f2533])).
fof(f2192, plain, (spl72_213 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl72_213])])).
fof(f2171, plain, (spl72_209 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl72_209])])).
fof(f2151, plain, (spl72_205 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl72_205])])).
fof(f2531, plain, (~ (h2(op1(e10, e12)) = op2(h2(e10), e21)) | ~ (h2(e10) = h2(op1(e11, e11))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), e21)) | ~ (h2(op1(e12, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e12, e11))) | ~ (h2(e11) = h2(op1(e12, e12))) | ~ (h2(op1(e12, e13)) = op2(e21, h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | ~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2530, f750])).
fof(f2530, plain, (~ (h2(e10) = h2(op1(e11, e11))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), e21)) | ~ (h2(op1(e12, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e12, e11))) | ~ (h2(e11) = h2(op1(e12, e12))) | ~ (h2(op1(e12, e13)) = op2(e21, h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | ~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2529, f2043])).
fof(f2529, plain, (~ (h2(op1(e11, e12)) = op2(h2(e11), e21)) | ~ (h2(op1(e12, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e12, e11))) | ~ (h2(e11) = h2(op1(e12, e12))) | ~ (h2(op1(e12, e13)) = op2(e21, h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | ~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2528, f750])).
fof(f2528, plain, (~ (h2(op1(e12, e10)) = op2(e21, h2(e10))) | ~ (h2(e13) = h2(op1(e12, e11))) | ~ (h2(e11) = h2(op1(e12, e12))) | ~ (h2(op1(e12, e13)) = op2(e21, h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | ~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2527, f750])).
fof(f2527, plain, (~ (h2(e13) = h2(op1(e12, e11))) | ~ (h2(e11) = h2(op1(e12, e12))) | ~ (h2(op1(e12, e13)) = op2(e21, h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | ~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2526, f2042])).
fof(f2526, plain, (~ (h2(op1(e12, e11)) = op2(e21, h2(e11))) | ~ (h2(e11) = h2(op1(e12, e12))) | ~ (h2(op1(e12, e13)) = op2(e21, h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | ~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2525, f750])).
fof(f2525, plain, (~ (h2(e11) = h2(op1(e12, e12))) | ~ (h2(op1(e12, e13)) = op2(e21, h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | ~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2524, f752])).
fof(f2524, plain, (~ (op2(e21, e21) = h2(op1(e12, e12))) | ~ (h2(op1(e12, e13)) = op2(e21, h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | ~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2523, f750])).
fof(f2523, plain, (~ (h2(op1(e12, e13)) = op2(e21, h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | ~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f2522, f750])).
fof(f2522, plain, (~ (h2(op1(e13, e12)) = op2(h2(e13), e21)) | ~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(forward_demodulation, [], [f817, f750])).
fof(f817, plain, (~ (e23 = h2(e13)) | sP65 | sP64 | sP63 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))), inference(cnf_transformation, [], [f95])).
fof(f2485, plain, (spl72_203 | spl72_201 | ~ spl72_248 | ~ spl72_249 | ~ spl72_250 | ~ spl72_251 | ~ spl72_252 | ~ spl72_253 | ~ spl72_254 | ~ spl72_255 | ~ spl72_256 | ~ spl72_257 | ~ spl72_258 | ~ spl72_259 | ~ spl72_260 | ~ spl72_261 | ~ spl72_262 | ~ spl72_263), inference(avatar_split_clause, [], [f2420, f2482, f2478, f2474, f2470, f2466, f2462, f2458, f2454, f2450, f2446, f2442, f2438, f2434, f2430, f2426, f2422, f2128, f2140])).
fof(f2140, plain, (spl72_203 <=> sP66), introduced(avatar_definition, [new_symbols(naming, [spl72_203])])).
fof(f2128, plain, (spl72_201 <=> sP67), introduced(avatar_definition, [new_symbols(naming, [spl72_201])])).
fof(f2420, plain, (~ (h1(e11) = h3(op1(e10, e10))) | ~ (h3(op1(e10, e11)) = op2(e20, h3(e11))) | ~ (op2(e20, e22) = h3(op1(e10, e12))) | ~ (op2(e20, e23) = h3(op1(e10, e13))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), e20)) | ~ (e20 = h3(op1(e11, e11))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66), inference(forward_demodulation, [], [f2419, f748])).
fof(f2419, plain, (~ (op2(e20, e20) = h3(op1(e10, e10))) | ~ (h3(op1(e10, e11)) = op2(e20, h3(e11))) | ~ (op2(e20, e22) = h3(op1(e10, e12))) | ~ (op2(e20, e23) = h3(op1(e10, e13))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), e20)) | ~ (e20 = h3(op1(e11, e11))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66), inference(forward_demodulation, [], [f2418, f2054])).
fof(f2418, plain, (~ (h3(op1(e10, e11)) = op2(e20, h3(e11))) | ~ (op2(e20, e22) = h3(op1(e10, e12))) | ~ (op2(e20, e23) = h3(op1(e10, e13))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), e20)) | ~ (e20 = h3(op1(e11, e11))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2417, f2054])).
fof(f2417, plain, (~ (op2(e20, e22) = h3(op1(e10, e12))) | ~ (op2(e20, e23) = h3(op1(e10, e13))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), e20)) | ~ (e20 = h3(op1(e11, e11))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2416, f2054])).
fof(f2416, plain, (~ (h3(op1(e10, e12)) = op2(h3(e10), e22)) | ~ (op2(e20, e23) = h3(op1(e10, e13))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), e20)) | ~ (e20 = h3(op1(e11, e11))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2415, f754])).
fof(f2415, plain, (~ (op2(e20, e23) = h3(op1(e10, e13))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), e20)) | ~ (e20 = h3(op1(e11, e11))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2414, f2054])).
fof(f2414, plain, (~ (h3(op1(e10, e13)) = op2(h3(e10), e23)) | ~ (h3(op1(e11, e10)) = op2(h3(e11), e20)) | ~ (e20 = h3(op1(e11, e11))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2413, f2044])).
fof(f2413, plain, (~ (h3(op1(e11, e10)) = op2(h3(e11), e20)) | ~ (e20 = h3(op1(e11, e11))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2412, f2054])).
fof(f2412, plain, (~ (e20 = h3(op1(e11, e11))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2411, f2051])).
fof(f2411, plain, (~ (h3(op1(e11, e12)) = op2(h3(e11), e22)) | ~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2410, f754])).
fof(f2410, plain, (~ (h3(op1(e11, e13)) = op2(h3(e11), e23)) | ~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2409, f2044])).
fof(f2409, plain, (~ (op2(e22, e20) = h3(op1(e12, e10))) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2408, f754])).
fof(f2408, plain, (~ (h3(op1(e12, e10)) = op2(h3(e12), e20)) | ~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2407, f2054])).
fof(f2407, plain, (~ (e23 = h3(op1(e12, e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2406, f2052])).
fof(f2406, plain, (~ (h3(op1(e12, e11)) = op2(e22, h3(e11))) | ~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2405, f754])).
fof(f2405, plain, (~ (h3(e11) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2404, f756])).
fof(f2404, plain, (~ (op2(e22, e22) = h3(op1(e12, e12))) | ~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2403, f754])).
fof(f2403, plain, (~ (op2(e22, e23) = h3(op1(e12, e13))) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2402, f754])).
fof(f2402, plain, (~ (h3(op1(e12, e13)) = op2(h3(e12), e23)) | ~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2401, f2044])).
fof(f2401, plain, (~ (op2(e23, e20) = h3(op1(e13, e10))) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2400, f2044])).
fof(f2400, plain, (~ (h3(op1(e13, e10)) = op2(h3(e13), e20)) | ~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2399, f2054])).
fof(f2399, plain, (~ (h3(op1(e13, e11)) = op2(e23, h3(e11))) | ~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2398, f2044])).
fof(f2398, plain, (~ (op2(e23, e22) = h3(op1(e13, e12))) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2397, f2044])).
fof(f2397, plain, (~ (h3(op1(e13, e12)) = op2(h3(e13), e22)) | ~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2396, f754])).
fof(f2396, plain, (~ (h4(e11) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2395, f760])).
fof(f2395, plain, (~ (op2(e23, e23) = h3(op1(e13, e13))) | sP67 | sP66 | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(forward_demodulation, [], [f2394, f2044])).
fof(f2394, plain, (sP67 | sP66 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(subsumption_resolution, [], [f2393, f2124])).
fof(f2124, plain, ~ sP68, inference(subsumption_resolution, [], [f776, f754])).
fof(f776, plain, (~ (e22 = h3(e12)) | ~ sP68), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP68), inference(nnf_transformation, [], [f91])).
fof(f2393, plain, (sP68 | sP67 | sP66 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(subsumption_resolution, [], [f821, f2044])).
fof(f821, plain, (~ (e23 = h3(e13)) | sP68 | sP67 | sP66 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))), inference(cnf_transformation, [], [f95])).
fof(f2362, plain, (~ spl72_229 | ~ spl72_230 | ~ spl72_231 | ~ spl72_232 | ~ spl72_233 | ~ spl72_234 | ~ spl72_235 | ~ spl72_236 | spl72_197 | spl72_193 | spl72_189 | ~ spl72_238 | ~ spl72_239 | ~ spl72_240 | ~ spl72_241 | ~ spl72_242 | ~ spl72_243 | ~ spl72_244 | ~ spl72_245), inference(avatar_split_clause, [], [f2361, f2347, f2343, f2339, f2335, f2331, f2327, f2323, f2319, f2064, f2084, f2104, f2311, f2307, f2303, f2299, f2295, f2291, f2287, f2283])).
fof(f2104, plain, (spl72_197 <=> sP69), introduced(avatar_definition, [new_symbols(naming, [spl72_197])])).
fof(f2084, plain, (spl72_193 <=> sP70), introduced(avatar_definition, [new_symbols(naming, [spl72_193])])).
fof(f2064, plain, (spl72_189 <=> sP71), introduced(avatar_definition, [new_symbols(naming, [spl72_189])])).
fof(f2361, plain, (~ (h4(op1(e10, e12)) = op2(h4(e10), e23)) | ~ (h4(e10) = h4(op1(e11, e11))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e12, e11))) | ~ (h4(e11) = h4(op1(e12, e12))) | ~ (h4(op1(e12, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2360, f758])).
fof(f2360, plain, (~ (h4(e10) = h4(op1(e11, e11))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e12, e11))) | ~ (h4(e11) = h4(op1(e12, e12))) | ~ (h4(op1(e12, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2359, f2062])).
fof(f2359, plain, (~ (h4(op1(e11, e12)) = op2(h4(e11), e23)) | ~ (h4(op1(e12, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e12, e11))) | ~ (h4(e11) = h4(op1(e12, e12))) | ~ (h4(op1(e12, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2358, f758])).
fof(f2358, plain, (~ (h4(op1(e12, e10)) = op2(e23, h4(e10))) | ~ (h4(e13) = h4(op1(e12, e11))) | ~ (h4(e11) = h4(op1(e12, e12))) | ~ (h4(op1(e12, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2357, f758])).
fof(f2357, plain, (~ (h4(e13) = h4(op1(e12, e11))) | ~ (h4(e11) = h4(op1(e12, e12))) | ~ (h4(op1(e12, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2356, f2061])).
fof(f2356, plain, (~ (h4(op1(e12, e11)) = op2(e23, h4(e11))) | ~ (h4(e11) = h4(op1(e12, e12))) | ~ (h4(op1(e12, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2355, f758])).
fof(f2355, plain, (~ (h4(e11) = h4(op1(e12, e12))) | ~ (h4(op1(e12, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2354, f760])).
fof(f2354, plain, (~ (op2(e23, e23) = h4(op1(e12, e12))) | ~ (h4(op1(e12, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2353, f758])).
fof(f2353, plain, (~ (h4(op1(e12, e13)) = op2(e23, h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2352, f758])).
fof(f2352, plain, (~ (h4(op1(e13, e12)) = op2(h4(e13), e23)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(forward_demodulation, [], [f2351, f758])).
fof(f2351, plain, (sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(subsumption_resolution, [], [f824, f758])).
fof(f824, plain, (~ (e23 = h4(e12)) | sP71 | sP70 | sP69 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))), inference(cnf_transformation, [], [f95])).
fof(f2261, plain, ~ spl72_225, inference(avatar_split_clause, [], [f2260, f2252])).
fof(f2260, plain, ~ sP60, inference(subsumption_resolution, [], [f808, f746])).
fof(f808, plain, (~ (e20 = h1(e12)) | ~ sP60), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP60), inference(nnf_transformation, [], [f83])).
fof(f2250, plain, (~ spl72_221 | ~ spl72_224), inference(avatar_split_clause, [], [f802, f2247, f2232])).
fof(f802, plain, (~ (e21 = h1(e10)) | ~ sP61), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP61), inference(nnf_transformation, [], [f84])).
fof(f2225, plain, (~ spl72_217 | ~ spl72_219), inference(avatar_split_clause, [], [f799, f2222, f2212])).
fof(f799, plain, (~ (e22 = h1(e11)) | ~ sP62), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP62), inference(nnf_transformation, [], [f85])).
fof(f2219, plain, (~ spl72_217 | ~ spl72_218), inference(avatar_split_clause, [], [f801, f2216, f2212])).
fof(f801, plain, (~ (e22 = h1(e13)) | ~ sP62), inference(cnf_transformation, [], [f165])).
fof(f2205, plain, (~ spl72_213 | ~ spl72_215), inference(avatar_split_clause, [], [f795, f2202, f2192])).
fof(f795, plain, (~ (e20 = h2(e11)) | ~ sP63), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP63), inference(nnf_transformation, [], [f86])).
fof(f2180, plain, ~ spl72_209, inference(avatar_split_clause, [], [f2179, f2171])).
fof(f2179, plain, ~ sP64, inference(subsumption_resolution, [], [f792, f750])).
fof(f792, plain, (~ (e21 = h2(e12)) | ~ sP64), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP64), inference(nnf_transformation, [], [f87])).
fof(f2169, plain, (~ spl72_205 | ~ spl72_208), inference(avatar_split_clause, [], [f786, f2166, f2151])).
fof(f786, plain, (~ (e22 = h2(e10)) | ~ sP65), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP65), inference(nnf_transformation, [], [f88])).
fof(f2149, plain, ~ spl72_203, inference(avatar_split_clause, [], [f2148, f2140])).
fof(f2148, plain, ~ sP66, inference(subsumption_resolution, [], [f782, f2054])).
fof(f782, plain, (~ (e20 = h3(e10)) | ~ sP66), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP66), inference(nnf_transformation, [], [f89])).
fof(f2135, plain, (~ spl72_201 | ~ spl72_202), inference(avatar_split_clause, [], [f779, f2132, f2128])).
fof(f779, plain, (~ (e21 = h3(e11)) | ~ sP67), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP67), inference(nnf_transformation, [], [f90])).
fof(f2122, plain, (~ spl72_197 | ~ spl72_200), inference(avatar_split_clause, [], [f770, f2119, f2104])).
fof(f770, plain, (~ (e20 = h4(e10)) | ~ sP69), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP69), inference(nnf_transformation, [], [f92])).
fof(f2117, plain, (~ spl72_197 | ~ spl72_199), inference(avatar_split_clause, [], [f771, f2114, f2104])).
fof(f771, plain, (~ (e20 = h4(e11)) | ~ sP69), inference(cnf_transformation, [], [f158])).
fof(f2097, plain, (~ spl72_193 | ~ spl72_195), inference(avatar_split_clause, [], [f767, f2094, f2084])).
fof(f767, plain, (~ (e21 = h4(e11)) | ~ sP70), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP70), inference(nnf_transformation, [], [f93])).
fof(f2091, plain, (~ spl72_193 | ~ spl72_194), inference(avatar_split_clause, [], [f769, f2088, f2084])).
fof(f769, plain, (~ (e21 = h4(e13)) | ~ sP70), inference(cnf_transformation, [], [f157])).
fof(f2082, plain, (~ spl72_189 | ~ spl72_192), inference(avatar_split_clause, [], [f762, f2079, f2064])).
fof(f762, plain, (~ (e22 = h4(e10)) | ~ sP71), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP71), inference(nnf_transformation, [], [f94])).
fof(f2071, plain, (~ spl72_189 | ~ spl72_190), inference(avatar_split_clause, [], [f765, f2068, f2064])).
fof(f765, plain, (~ (e22 = h4(e13)) | ~ sP71), inference(cnf_transformation, [], [f156])).
fof(f2027, plain, spl72_86, inference(avatar_split_clause, [], [f744, f1252])).
fof(f744, plain, (e21 = op2(e22, e22)), inference(cnf_transformation, [], [f13])).
fof(f2026, plain, spl72_22, inference(avatar_split_clause, [], [f741, f948])).
fof(f741, plain, (e11 = op1(e12, e12)), inference(cnf_transformation, [], [f12])).
fof(f2025, plain, (spl72_188 | spl72_187 | spl72_186 | spl72_185 | spl72_184 | spl72_183 | spl72_182 | spl72_181 | spl72_180 | spl72_179 | spl72_178 | spl72_177 | spl72_176 | spl72_175 | spl72_174 | spl72_68), inference(avatar_split_clause, [], [f729, f1175, f1897, f1905, f1913, f1921, f1929, f1937, f1945, f1953, f1961, f1969, f1977, f1985, f1993, f2001, f2009])).
fof(f2009, plain, (spl72_188 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl72_188])])).
fof(f2001, plain, (spl72_187 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl72_187])])).
fof(f1993, plain, (spl72_186 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl72_186])])).
fof(f1985, plain, (spl72_185 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl72_185])])).
fof(f1977, plain, (spl72_184 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl72_184])])).
fof(f1969, plain, (spl72_183 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl72_183])])).
fof(f1961, plain, (spl72_182 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl72_182])])).
fof(f1953, plain, (spl72_181 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl72_181])])).
fof(f1945, plain, (spl72_180 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl72_180])])).
fof(f1937, plain, (spl72_179 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl72_179])])).
fof(f1929, plain, (spl72_178 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl72_178])])).
fof(f1921, plain, (spl72_177 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl72_177])])).
fof(f1913, plain, (spl72_176 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl72_176])])).
fof(f1905, plain, (spl72_175 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl72_175])])).
fof(f1897, plain, (spl72_174 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl72_174])])).
fof(f729, plain, ((e23 = op2(e23, e23)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(cnf_transformation, [], [f82])).
fof(f82, plain, ((((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45) & ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e23, e23))) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30)), inference(definition_folding, [], [f11, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52])).
fof(f52, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20))) | ~ sP30), inference(usedef, [], [e52])).
fof(e52, plain, (sP30 <=> (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f53, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (op2(e20, e20) = e21)) | ~ sP31), inference(usedef, [], [e53])).
fof(e53, plain, (sP31 <=> (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (op2(e20, e20) = e21))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f54, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (op2(e20, e20) = e22)) | ~ sP32), inference(usedef, [], [e54])).
fof(e54, plain, (sP32 <=> (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (op2(e20, e20) = e22))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f55, plain, ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (op2(e20, e20) = e23)) | ~ sP33), inference(usedef, [], [e55])).
fof(e55, plain, (sP33 <=> (((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (op2(e20, e20) = e23))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f56, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e21, e21))) | ~ sP34), inference(usedef, [], [e56])).
fof(e56, plain, (sP34 <=> (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f57, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e21, e21))) | ~ sP35), inference(usedef, [], [e57])).
fof(e57, plain, (sP35 <=> (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f58, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e21, e21))) | ~ sP36), inference(usedef, [], [e58])).
fof(e58, plain, (sP36 <=> (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f59, plain, ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e21, e21))) | ~ sP37), inference(usedef, [], [e59])).
fof(e59, plain, (sP37 <=> (((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f60, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e22, e22))) | ~ sP38), inference(usedef, [], [e60])).
fof(e60, plain, (sP38 <=> (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f61, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e22, e22))) | ~ sP39), inference(usedef, [], [e61])).
fof(e61, plain, (sP39 <=> (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f62, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e22, e22))) | ~ sP40), inference(usedef, [], [e62])).
fof(e62, plain, (sP40 <=> (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f63, plain, ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e22, e22))) | ~ sP41), inference(usedef, [], [e63])).
fof(e63, plain, (sP41 <=> (((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f64, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e23, e23))) | ~ sP42), inference(usedef, [], [e64])).
fof(e64, plain, (sP42 <=> (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f65, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e23, e23))) | ~ sP43), inference(usedef, [], [e65])).
fof(e65, plain, (sP43 <=> (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f66, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e23, e23))) | ~ sP44), inference(usedef, [], [e66])).
fof(e66, plain, (sP44 <=> (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e23, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f67, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP45), inference(usedef, [], [e67])).
fof(e67, plain, (sP45 <=> ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f68, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP46), inference(usedef, [], [e68])).
fof(e68, plain, (sP46 <=> ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f69, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP47), inference(usedef, [], [e69])).
fof(e69, plain, (sP47 <=> ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f70, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP48), inference(usedef, [], [e70])).
fof(e70, plain, (sP48 <=> ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f71, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP49), inference(usedef, [], [e71])).
fof(e71, plain, (sP49 <=> ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f72, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP50), inference(usedef, [], [e72])).
fof(e72, plain, (sP50 <=> ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f73, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP51), inference(usedef, [], [e73])).
fof(e73, plain, (sP51 <=> ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f74, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP52), inference(usedef, [], [e74])).
fof(e74, plain, (sP52 <=> ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f75, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP53), inference(usedef, [], [e75])).
fof(e75, plain, (sP53 <=> ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f76, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP54), inference(usedef, [], [e76])).
fof(e76, plain, (sP54 <=> ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f77, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP55), inference(usedef, [], [e77])).
fof(e77, plain, (sP55 <=> ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f78, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP56), inference(usedef, [], [e78])).
fof(e78, plain, (sP56 <=> ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f79, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP57), inference(usedef, [], [e79])).
fof(e79, plain, (sP57 <=> ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f80, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP58), inference(usedef, [], [e80])).
fof(e80, plain, (sP58 <=> ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f81, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP59), inference(usedef, [], [e81])).
fof(e81, plain, (sP59 <=> ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f11, plain, ((((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e23 = op2(e23, e23)) & (e23 = op2(e23, e23))) | ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20)))) & ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e23, e23))) | (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e23, e23))) | (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e23, e23))) | (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e23, e23))) | (((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e22, e22))) | (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e22, e22))) | (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e22, e22))) | (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e22, e22))) | (((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e21, e21))) | (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e21, e21))) | (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e21, e21))) | (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e21, e21))) | (((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (op2(e20, e20) = e23)) | (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (op2(e20, e20) = e22)) | (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (op2(e20, e20) = e21)) | (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax11)).
fof(f2024, plain, (spl72_188 | spl72_187 | spl72_186 | spl72_185 | spl72_184 | spl72_183 | spl72_182 | spl72_181 | spl72_180 | spl72_179 | spl72_178 | spl72_177 | spl72_176 | spl72_175 | spl72_174 | ~ spl72_128 | spl72_113), inference(avatar_split_clause, [], [f730, f1367, f1430, f1897, f1905, f1913, f1921, f1929, f1937, f1945, f1953, f1961, f1969, f1977, f1985, f1993, f2001, f2009])).
fof(f730, plain, ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30), inference(cnf_transformation, [], [f82])).
fof(f2021, plain, (spl72_173 | spl72_172 | spl72_171 | spl72_170 | spl72_169 | spl72_168 | spl72_167 | spl72_166 | spl72_165 | spl72_164 | spl72_163 | spl72_162 | spl72_161 | spl72_160 | spl72_159 | spl72_68), inference(avatar_split_clause, [], [f734, f1175, f1747, f1757, f1767, f1777, f1787, f1797, f1807, f1817, f1827, f1837, f1847, f1857, f1867, f1877, f1887])).
fof(f1887, plain, (spl72_173 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl72_173])])).
fof(f1877, plain, (spl72_172 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl72_172])])).
fof(f1867, plain, (spl72_171 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl72_171])])).
fof(f1857, plain, (spl72_170 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl72_170])])).
fof(f1847, plain, (spl72_169 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl72_169])])).
fof(f1837, plain, (spl72_168 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl72_168])])).
fof(f1827, plain, (spl72_167 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl72_167])])).
fof(f1817, plain, (spl72_166 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl72_166])])).
fof(f1807, plain, (spl72_165 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl72_165])])).
fof(f1797, plain, (spl72_164 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl72_164])])).
fof(f1787, plain, (spl72_163 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl72_163])])).
fof(f1777, plain, (spl72_162 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl72_162])])).
fof(f1767, plain, (spl72_161 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl72_161])])).
fof(f1757, plain, (spl72_160 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl72_160])])).
fof(f1747, plain, (spl72_159 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl72_159])])).
fof(f734, plain, ((e23 = op2(e23, e23)) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45), inference(cnf_transformation, [], [f82])).
fof(f2017, plain, (spl72_173 | spl72_172 | spl72_171 | spl72_170 | spl72_169 | spl72_168 | spl72_167 | spl72_166 | spl72_165 | spl72_164 | spl72_163 | spl72_162 | spl72_161 | spl72_160 | spl72_159 | ~ spl72_72 | ~ spl72_67), inference(avatar_split_clause, [], [f738, f1171, f1192, f1747, f1757, f1767, f1777, f1787, f1797, f1807, f1817, f1827, f1837, f1847, f1857, f1867, f1877, f1887])).
fof(f738, plain, (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22)) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45), inference(cnf_transformation, [], [f82])).
fof(f2016, plain, (spl72_173 | spl72_172 | spl72_171 | spl72_170 | spl72_169 | spl72_168 | spl72_167 | spl72_166 | spl72_165 | spl72_164 | spl72_163 | spl72_162 | spl72_161 | spl72_160 | spl72_159 | ~ spl72_68), inference(avatar_split_clause, [], [f826, f1175, f1747, f1757, f1767, f1777, f1787, f1797, f1807, f1817, f1827, f1837, f1847, f1857, f1867, f1877, f1887])).
fof(f826, plain, (~ (e23 = op2(e23, e23)) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45), inference(duplicate_literal_removal, [], [f739])).
fof(f739, plain, (~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23)) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45), inference(cnf_transformation, [], [f82])).
fof(f2015, plain, (~ spl72_188 | spl72_125), inference(avatar_split_clause, [], [f724, f1418, f2009])).
fof(f724, plain, ((e20 = op2(e20, e20)) | ~ sP30), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20))) | ~ sP30), inference(nnf_transformation, [], [f52])).
fof(f2005, plain, (~ spl72_187 | ~ spl72_86 | spl72_91), inference(avatar_split_clause, [], [f722, f1273, f1252, f2001])).
fof(f722, plain, ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22)) | ~ sP31), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (op2(e20, e20) = e21)) | ~ sP31), inference(nnf_transformation, [], [f53])).
fof(f1999, plain, (~ spl72_186 | spl72_127), inference(avatar_split_clause, [], [f714, f1426, f1993])).
fof(f714, plain, ((op2(e20, e20) = e22) | ~ sP32), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (op2(e20, e20) = e22)) | ~ sP32), inference(nnf_transformation, [], [f54])).
fof(f1991, plain, (~ spl72_185 | spl72_128), inference(avatar_split_clause, [], [f709, f1430, f1985])).
fof(f709, plain, ((op2(e20, e20) = e23) | ~ sP33), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (op2(e20, e20) = e23)) | ~ sP33), inference(nnf_transformation, [], [f55])).
fof(f1990, plain, (~ spl72_185 | ~ spl72_128 | spl72_113), inference(avatar_split_clause, [], [f710, f1367, f1430, f1985])).
fof(f710, plain, ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23) | ~ sP33), inference(cnf_transformation, [], [f152])).
fof(f1982, plain, (~ spl72_184 | ~ spl72_105 | spl72_110), inference(avatar_split_clause, [], [f706, f1354, f1333, f1977])).
fof(f706, plain, ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21)) | ~ sP34), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e21, e21))) | ~ sP34), inference(nnf_transformation, [], [f56])).
fof(f1980, plain, (~ spl72_184 | ~ spl72_65 | spl72_80), inference(avatar_split_clause, [], [f708, f1226, f1163, f1977])).
fof(f708, plain, ((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23)) | ~ sP34), inference(cnf_transformation, [], [f151])).
fof(f1975, plain, (~ spl72_183 | spl72_106), inference(avatar_split_clause, [], [f699, f1337, f1969])).
fof(f699, plain, ((e21 = op2(e21, e21)) | ~ sP35), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e21, e21))) | ~ sP35), inference(nnf_transformation, [], [f57])).
fof(f1967, plain, (~ spl72_182 | spl72_107), inference(avatar_split_clause, [], [f694, f1341, f1961])).
fof(f694, plain, ((e22 = op2(e21, e21)) | ~ sP36), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e21, e21))) | ~ sP36), inference(nnf_transformation, [], [f58])).
fof(f1959, plain, (~ spl72_181 | spl72_108), inference(avatar_split_clause, [], [f689, f1345, f1953])).
fof(f689, plain, ((e23 = op2(e21, e21)) | ~ sP37), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e21, e21))) | ~ sP37), inference(nnf_transformation, [], [f59])).
fof(f1951, plain, (~ spl72_180 | spl72_85), inference(avatar_split_clause, [], [f684, f1248, f1945])).
fof(f684, plain, ((e20 = op2(e22, e22)) | ~ sP38), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e22, e22))) | ~ sP38), inference(nnf_transformation, [], [f60])).
fof(f1941, plain, (~ spl72_179 | ~ spl72_86 | spl72_91), inference(avatar_split_clause, [], [f682, f1273, f1252, f1937])).
fof(f682, plain, ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22)) | ~ sP39), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e22, e22))) | ~ sP39), inference(nnf_transformation, [], [f61])).
fof(f1935, plain, (~ spl72_178 | spl72_87), inference(avatar_split_clause, [], [f674, f1256, f1929])).
fof(f674, plain, ((e22 = op2(e22, e22)) | ~ sP40), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e22, e22))) | ~ sP40), inference(nnf_transformation, [], [f62])).
fof(f1927, plain, (~ spl72_177 | spl72_88), inference(avatar_split_clause, [], [f669, f1260, f1921])).
fof(f669, plain, ((e23 = op2(e22, e22)) | ~ sP41), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)) & (e23 = op2(e22, e22))) | ~ sP41), inference(nnf_transformation, [], [f63])).
fof(f1918, plain, (~ spl72_176 | ~ spl72_105 | spl72_110), inference(avatar_split_clause, [], [f666, f1354, f1333, f1913])).
fof(f666, plain, ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21)) | ~ sP42), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e23, e23))) | ~ sP42), inference(nnf_transformation, [], [f64])).
fof(f1916, plain, (~ spl72_176 | ~ spl72_65 | spl72_80), inference(avatar_split_clause, [], [f668, f1226, f1163, f1913])).
fof(f668, plain, ((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23)) | ~ sP42), inference(cnf_transformation, [], [f143])).
fof(f1909, plain, (~ spl72_175 | ~ spl72_86 | spl72_91), inference(avatar_split_clause, [], [f662, f1273, f1252, f1905])).
fof(f662, plain, ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22)) | ~ sP43), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)) & (e21 = op2(e23, e23))) | ~ sP43), inference(nnf_transformation, [], [f65])).
fof(f1903, plain, (~ spl72_174 | spl72_67), inference(avatar_split_clause, [], [f654, f1171, f1897])).
fof(f654, plain, ((e22 = op2(e23, e23)) | ~ sP44), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)) & (e22 = op2(e23, e23))) | ~ sP44), inference(nnf_transformation, [], [f66])).
fof(f1900, plain, (~ spl72_174 | ~ spl72_67 | spl72_72), inference(avatar_split_clause, [], [f658, f1192, f1171, f1897])).
fof(f658, plain, ((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23)) | ~ sP44), inference(cnf_transformation, [], [f141])).
fof(f1895, plain, (~ spl72_173 | spl72_125), inference(avatar_split_clause, [], [f648, f1418, f1887])).
fof(f648, plain, ((e20 = op2(e20, e20)) | ~ sP45), inference(cnf_transformation, [], [f140])).
fof(f140, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e20 = op2(e20, e20)) & (e20 = op2(e20, e20))) | ~ sP45), inference(nnf_transformation, [], [f67])).
fof(f1893, plain, (~ spl72_173 | ~ spl72_125), inference(avatar_split_clause, [], [f827, f1418, f1887])).
fof(f827, plain, (~ (e20 = op2(e20, e20)) | ~ sP45), inference(duplicate_literal_removal, [], [f650])).
fof(f650, plain, (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)) | ~ sP45), inference(cnf_transformation, [], [f140])).
fof(f1885, plain, (~ spl72_172 | spl72_121), inference(avatar_split_clause, [], [f642, f1401, f1877])).
fof(f642, plain, ((e20 = op2(e20, e21)) | ~ sP46), inference(cnf_transformation, [], [f139])).
fof(f139, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e20 = op2(e21, e20)) & (e20 = op2(e20, e21))) | ~ sP46), inference(nnf_transformation, [], [f68])).
fof(f1875, plain, (~ spl72_171 | spl72_117), inference(avatar_split_clause, [], [f636, f1384, f1867])).
fof(f636, plain, ((e20 = op2(e20, e22)) | ~ sP47), inference(cnf_transformation, [], [f138])).
fof(f138, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e20 = op2(e22, e20)) & (e20 = op2(e20, e22))) | ~ sP47), inference(nnf_transformation, [], [f69])).
fof(f1874, plain, (~ spl72_171 | spl72_93), inference(avatar_split_clause, [], [f637, f1282, f1867])).
fof(f637, plain, ((e20 = op2(e22, e20)) | ~ sP47), inference(cnf_transformation, [], [f138])).
fof(f1865, plain, (~ spl72_170 | spl72_113), inference(avatar_split_clause, [], [f630, f1367, f1857])).
fof(f630, plain, ((e20 = op2(e20, e23)) | ~ sP48), inference(cnf_transformation, [], [f137])).
fof(f137, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e20 = op2(e23, e20)) & (e20 = op2(e20, e23))) | ~ sP48), inference(nnf_transformation, [], [f70])).
fof(f1864, plain, (~ spl72_170 | spl72_77), inference(avatar_split_clause, [], [f631, f1214, f1857])).
fof(f631, plain, ((e20 = op2(e23, e20)) | ~ sP48), inference(cnf_transformation, [], [f137])).
fof(f1860, plain, (~ spl72_170 | ~ spl72_68), inference(avatar_split_clause, [], [f830, f1175, f1857])).
fof(f830, plain, (~ (e23 = op2(e23, e23)) | ~ sP48), inference(duplicate_literal_removal, [], [f635])).
fof(f635, plain, (~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23)) | ~ sP48), inference(cnf_transformation, [], [f137])).
fof(f1855, plain, (~ spl72_169 | spl72_110), inference(avatar_split_clause, [], [f624, f1354, f1847])).
fof(f624, plain, ((e21 = op2(e21, e20)) | ~ sP49), inference(cnf_transformation, [], [f136])).
fof(f136, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e21 = op2(e20, e21)) & (e21 = op2(e21, e20))) | ~ sP49), inference(nnf_transformation, [], [f71])).
fof(f1854, plain, (~ spl72_169 | spl72_122), inference(avatar_split_clause, [], [f625, f1405, f1847])).
fof(f625, plain, ((e21 = op2(e20, e21)) | ~ sP49), inference(cnf_transformation, [], [f136])).
fof(f1853, plain, (~ spl72_169 | ~ spl72_125), inference(avatar_split_clause, [], [f831, f1418, f1847])).
fof(f831, plain, (~ (e20 = op2(e20, e20)) | ~ sP49), inference(duplicate_literal_removal, [], [f626])).
fof(f626, plain, (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)) | ~ sP49), inference(cnf_transformation, [], [f136])).
fof(f1845, plain, (~ spl72_168 | spl72_106), inference(avatar_split_clause, [], [f618, f1337, f1837])).
fof(f618, plain, ((e21 = op2(e21, e21)) | ~ sP50), inference(cnf_transformation, [], [f135])).
fof(f135, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e21 = op2(e21, e21)) & (e21 = op2(e21, e21))) | ~ sP50), inference(nnf_transformation, [], [f72])).
fof(f1835, plain, (~ spl72_167 | spl72_102), inference(avatar_split_clause, [], [f612, f1320, f1827])).
fof(f612, plain, ((e21 = op2(e21, e22)) | ~ sP51), inference(cnf_transformation, [], [f134])).
fof(f134, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e21 = op2(e22, e21)) & (e21 = op2(e21, e22))) | ~ sP51), inference(nnf_transformation, [], [f73])).
fof(f1825, plain, (~ spl72_166 | spl72_98), inference(avatar_split_clause, [], [f606, f1303, f1817])).
fof(f606, plain, ((e21 = op2(e21, e23)) | ~ sP52), inference(cnf_transformation, [], [f133])).
fof(f133, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e21 = op2(e23, e21)) & (e21 = op2(e21, e23))) | ~ sP52), inference(nnf_transformation, [], [f74])).
fof(f1824, plain, (~ spl72_166 | spl72_74), inference(avatar_split_clause, [], [f607, f1201, f1817])).
fof(f607, plain, ((e21 = op2(e23, e21)) | ~ sP52), inference(cnf_transformation, [], [f133])).
fof(f1821, plain, (~ spl72_166 | ~ spl72_72 | ~ spl72_67), inference(avatar_split_clause, [], [f610, f1171, f1192, f1817])).
fof(f610, plain, (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22)) | ~ sP52), inference(cnf_transformation, [], [f133])).
fof(f1820, plain, (~ spl72_166 | ~ spl72_68), inference(avatar_split_clause, [], [f834, f1175, f1817])).
fof(f834, plain, (~ (e23 = op2(e23, e23)) | ~ sP52), inference(duplicate_literal_removal, [], [f611])).
fof(f611, plain, (~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23)) | ~ sP52), inference(cnf_transformation, [], [f133])).
fof(f1815, plain, (~ spl72_165 | spl72_95), inference(avatar_split_clause, [], [f600, f1290, f1807])).
fof(f600, plain, ((e22 = op2(e22, e20)) | ~ sP53), inference(cnf_transformation, [], [f132])).
fof(f132, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e22 = op2(e20, e22)) & (e22 = op2(e22, e20))) | ~ sP53), inference(nnf_transformation, [], [f75])).
fof(f1814, plain, (~ spl72_165 | spl72_119), inference(avatar_split_clause, [], [f601, f1392, f1807])).
fof(f601, plain, ((e22 = op2(e20, e22)) | ~ sP53), inference(cnf_transformation, [], [f132])).
fof(f1813, plain, (~ spl72_165 | ~ spl72_125), inference(avatar_split_clause, [], [f835, f1418, f1807])).
fof(f835, plain, (~ (e20 = op2(e20, e20)) | ~ sP53), inference(duplicate_literal_removal, [], [f602])).
fof(f602, plain, (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)) | ~ sP53), inference(cnf_transformation, [], [f132])).
fof(f1805, plain, (~ spl72_164 | spl72_91), inference(avatar_split_clause, [], [f594, f1273, f1797])).
fof(f594, plain, ((e22 = op2(e22, e21)) | ~ sP54), inference(cnf_transformation, [], [f131])).
fof(f131, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e22 = op2(e21, e22)) & (e22 = op2(e22, e21))) | ~ sP54), inference(nnf_transformation, [], [f76])).
fof(f1795, plain, (~ spl72_163 | spl72_87), inference(avatar_split_clause, [], [f588, f1256, f1787])).
fof(f588, plain, ((e22 = op2(e22, e22)) | ~ sP55), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e22 = op2(e22, e22)) & (e22 = op2(e22, e22))) | ~ sP55), inference(nnf_transformation, [], [f77])).
fof(f1785, plain, (~ spl72_162 | spl72_83), inference(avatar_split_clause, [], [f582, f1239, f1777])).
fof(f582, plain, ((e22 = op2(e22, e23)) | ~ sP56), inference(cnf_transformation, [], [f129])).
fof(f129, plain, (((~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & (~ (e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & (~ (e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & (~ (e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & (e22 = op2(e23, e22)) & (e22 = op2(e22, e23))) | ~ sP56), inference(nnf_transformation, [], [f78])).
fof(f1784, plain, (~ spl72_162 | spl72_71), inference(avatar_split_clause, [], [f583, f1188, f1777])).
fof(f583, plain, ((e22 = op2(e23, e22)) | ~ sP56), inference(cnf_transformation, [], [f129])).
fof(f1780, plain, (~ spl72_162 | ~ spl72_68), inference(avatar_split_clause, [], [f838, f1175, f1777])).
fof(f838, plain, (~ (e23 = op2(e23, e23)) | ~ sP56), inference(duplicate_literal_removal, [], [f587])).
fof(f587, plain, (~ (e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23)) | ~ sP56), inference(cnf_transformation, [], [f129])).
fof(f1775, plain, (~ spl72_161 | spl72_80), inference(avatar_split_clause, [], [f576, f1226, f1767])).
fof(f576, plain, ((e23 = op2(e23, e20)) | ~ sP57), inference(cnf_transformation, [], [f128])).
fof(f128, plain, (((~ (op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & (~ (op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & (~ (op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))) & (e23 = op2(e20, e23)) & (e23 = op2(e23, e20))) | ~ sP57), inference(nnf_transformation, [], [f79])).
fof(f1773, plain, (~ spl72_161 | ~ spl72_125), inference(avatar_split_clause, [], [f839, f1418, f1767])).
fof(f839, plain, (~ (e20 = op2(e20, e20)) | ~ sP57), inference(duplicate_literal_removal, [], [f578])).
fof(f578, plain, (~ (e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)) | ~ sP57), inference(cnf_transformation, [], [f128])).
fof(f1765, plain, (~ spl72_160 | spl72_76), inference(avatar_split_clause, [], [f570, f1209, f1757])).
fof(f570, plain, ((e23 = op2(e23, e21)) | ~ sP58), inference(cnf_transformation, [], [f127])).
fof(f127, plain, (((~ (e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & (~ (e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & (~ (e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & (~ (e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & (e23 = op2(e21, e23)) & (e23 = op2(e23, e21))) | ~ sP58), inference(nnf_transformation, [], [f80])).
fof(f1754, plain, (~ spl72_159 | spl72_84), inference(avatar_split_clause, [], [f565, f1243, f1747])).
fof(f565, plain, ((e23 = op2(e22, e23)) | ~ sP59), inference(cnf_transformation, [], [f126])).
fof(f126, plain, (((~ (e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & (~ (e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & (~ (e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & (~ (e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & (e23 = op2(e22, e23)) & (e23 = op2(e23, e22))) | ~ sP59), inference(nnf_transformation, [], [f81])).
fof(f1745, plain, (spl72_158 | spl72_157 | spl72_156 | spl72_155 | spl72_154 | spl72_153 | spl72_152 | spl72_151 | spl72_150 | spl72_149 | spl72_148 | spl72_147 | spl72_146 | spl72_145 | spl72_144 | spl72_4), inference(avatar_split_clause, [], [f553, f871, f1617, f1625, f1633, f1641, f1649, f1657, f1665, f1673, f1681, f1689, f1697, f1705, f1713, f1721, f1729])).
fof(f1729, plain, (spl72_158 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl72_158])])).
fof(f1721, plain, (spl72_157 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl72_157])])).
fof(f1713, plain, (spl72_156 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl72_156])])).
fof(f1705, plain, (spl72_155 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl72_155])])).
fof(f1697, plain, (spl72_154 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl72_154])])).
fof(f1689, plain, (spl72_153 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl72_153])])).
fof(f1681, plain, (spl72_152 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl72_152])])).
fof(f1673, plain, (spl72_151 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl72_151])])).
fof(f1665, plain, (spl72_150 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl72_150])])).
fof(f1657, plain, (spl72_149 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl72_149])])).
fof(f1649, plain, (spl72_148 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl72_148])])).
fof(f1641, plain, (spl72_147 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl72_147])])).
fof(f1633, plain, (spl72_146 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl72_146])])).
fof(f1625, plain, (spl72_145 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl72_145])])).
fof(f1617, plain, (spl72_144 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl72_144])])).
fof(f553, plain, ((e13 = op1(e13, e13)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f51])).
fof(f51, plain, ((((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15) & ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e13, e13))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f10, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21])).
fof(f21, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10))) | ~ sP0), inference(usedef, [], [e21])).
fof(e21, plain, (sP0 <=> (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f22, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (op1(e10, e10) = e11)) | ~ sP1), inference(usedef, [], [e22])).
fof(e22, plain, (sP1 <=> (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (op1(e10, e10) = e11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f23, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (op1(e10, e10) = e12)) | ~ sP2), inference(usedef, [], [e23])).
fof(e23, plain, (sP2 <=> (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (op1(e10, e10) = e12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f24, plain, ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (op1(e10, e10) = e13)) | ~ sP3), inference(usedef, [], [e24])).
fof(e24, plain, (sP3 <=> (((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (op1(e10, e10) = e13))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f25, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e11, e11))) | ~ sP4), inference(usedef, [], [e25])).
fof(e25, plain, (sP4 <=> (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f26, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e11, e11))) | ~ sP5), inference(usedef, [], [e26])).
fof(e26, plain, (sP5 <=> (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f27, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e11, e11))) | ~ sP6), inference(usedef, [], [e27])).
fof(e27, plain, (sP6 <=> (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f28, plain, ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e11, e11))) | ~ sP7), inference(usedef, [], [e28])).
fof(e28, plain, (sP7 <=> (((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f29, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e12, e12))) | ~ sP8), inference(usedef, [], [e29])).
fof(e29, plain, (sP8 <=> (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f30, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e12, e12))) | ~ sP9), inference(usedef, [], [e30])).
fof(e30, plain, (sP9 <=> (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f31, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e12, e12))) | ~ sP10), inference(usedef, [], [e31])).
fof(e31, plain, (sP10 <=> (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f32, plain, ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e12, e12))) | ~ sP11), inference(usedef, [], [e32])).
fof(e32, plain, (sP11 <=> (((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f33, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e13, e13))) | ~ sP12), inference(usedef, [], [e33])).
fof(e33, plain, (sP12 <=> (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f34, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e13, e13))) | ~ sP13), inference(usedef, [], [e34])).
fof(e34, plain, (sP13 <=> (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f35, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e13, e13))) | ~ sP14), inference(usedef, [], [e35])).
fof(e35, plain, (sP14 <=> (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e13, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f36, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP15), inference(usedef, [], [e36])).
fof(e36, plain, (sP15 <=> ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f37, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP16), inference(usedef, [], [e37])).
fof(e37, plain, (sP16 <=> ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f38, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP17), inference(usedef, [], [e38])).
fof(e38, plain, (sP17 <=> ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f39, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP18), inference(usedef, [], [e39])).
fof(e39, plain, (sP18 <=> ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f40, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP19), inference(usedef, [], [e40])).
fof(e40, plain, (sP19 <=> ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f41, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP20), inference(usedef, [], [e41])).
fof(e41, plain, (sP20 <=> ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f42, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP21), inference(usedef, [], [e42])).
fof(e42, plain, (sP21 <=> ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f43, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP22), inference(usedef, [], [e43])).
fof(e43, plain, (sP22 <=> ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f44, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP23), inference(usedef, [], [e44])).
fof(e44, plain, (sP23 <=> ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f45, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP24), inference(usedef, [], [e45])).
fof(e45, plain, (sP24 <=> ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f46, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP25), inference(usedef, [], [e46])).
fof(e46, plain, (sP25 <=> ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f47, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP26), inference(usedef, [], [e47])).
fof(e47, plain, (sP26 <=> ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f48, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP27), inference(usedef, [], [e48])).
fof(e48, plain, (sP27 <=> ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f49, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP28), inference(usedef, [], [e49])).
fof(e49, plain, (sP28 <=> ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f50, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP29), inference(usedef, [], [e50])).
fof(e50, plain, (sP29 <=> ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f10, plain, ((((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e13 = op1(e13, e13)) & (e13 = op1(e13, e13))) | ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10)))) & ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e13, e13))) | (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e13, e13))) | (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e13, e13))) | (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e13, e13))) | (((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e12, e12))) | (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e12, e12))) | (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e12, e12))) | (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e12, e12))) | (((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e11, e11))) | (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e11, e11))) | (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e11, e11))) | (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e11, e11))) | (((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (op1(e10, e10) = e13)) | (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (op1(e10, e10) = e12)) | (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (op1(e10, e10) = e11)) | (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax10)).
fof(f1744, plain, (spl72_158 | spl72_157 | spl72_156 | spl72_155 | spl72_154 | spl72_153 | spl72_152 | spl72_151 | spl72_150 | spl72_149 | spl72_148 | spl72_147 | spl72_146 | spl72_145 | spl72_144 | ~ spl72_64 | spl72_49), inference(avatar_split_clause, [], [f554, f1063, f1126, f1617, f1625, f1633, f1641, f1649, f1657, f1665, f1673, f1681, f1689, f1697, f1705, f1713, f1721, f1729])).
fof(f554, plain, ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f51])).
fof(f1741, plain, (spl72_143 | spl72_142 | spl72_141 | spl72_140 | spl72_139 | spl72_138 | spl72_137 | spl72_136 | spl72_135 | spl72_134 | spl72_133 | spl72_132 | spl72_131 | spl72_130 | spl72_129 | spl72_4), inference(avatar_split_clause, [], [f558, f871, f1467, f1477, f1487, f1497, f1507, f1517, f1527, f1537, f1547, f1557, f1567, f1577, f1587, f1597, f1607])).
fof(f1607, plain, (spl72_143 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl72_143])])).
fof(f1597, plain, (spl72_142 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl72_142])])).
fof(f1587, plain, (spl72_141 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl72_141])])).
fof(f1577, plain, (spl72_140 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl72_140])])).
fof(f1567, plain, (spl72_139 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl72_139])])).
fof(f1557, plain, (spl72_138 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl72_138])])).
fof(f1547, plain, (spl72_137 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl72_137])])).
fof(f1537, plain, (spl72_136 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl72_136])])).
fof(f1527, plain, (spl72_135 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl72_135])])).
fof(f1517, plain, (spl72_134 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl72_134])])).
fof(f1507, plain, (spl72_133 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl72_133])])).
fof(f1497, plain, (spl72_132 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl72_132])])).
fof(f1487, plain, (spl72_131 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl72_131])])).
fof(f1477, plain, (spl72_130 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl72_130])])).
fof(f1467, plain, (spl72_129 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl72_129])])).
fof(f558, plain, ((e13 = op1(e13, e13)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f51])).
fof(f1737, plain, (spl72_143 | spl72_142 | spl72_141 | spl72_140 | spl72_139 | spl72_138 | spl72_137 | spl72_136 | spl72_135 | spl72_134 | spl72_133 | spl72_132 | spl72_131 | spl72_130 | spl72_129 | ~ spl72_8 | ~ spl72_3), inference(avatar_split_clause, [], [f562, f867, f888, f1467, f1477, f1487, f1497, f1507, f1517, f1527, f1537, f1547, f1557, f1567, f1577, f1587, f1597, f1607])).
fof(f562, plain, (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f51])).
fof(f1736, plain, (spl72_143 | spl72_142 | spl72_141 | spl72_140 | spl72_139 | spl72_138 | spl72_137 | spl72_136 | spl72_135 | spl72_134 | spl72_133 | spl72_132 | spl72_131 | spl72_130 | spl72_129 | ~ spl72_4), inference(avatar_split_clause, [], [f842, f871, f1467, f1477, f1487, f1497, f1507, f1517, f1527, f1537, f1547, f1557, f1567, f1577, f1587, f1597, f1607])).
fof(f842, plain, (~ (e13 = op1(e13, e13)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(duplicate_literal_removal, [], [f563])).
fof(f563, plain, (~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f51])).
fof(f1735, plain, (~ spl72_158 | spl72_61), inference(avatar_split_clause, [], [f548, f1114, f1729])).
fof(f548, plain, ((e10 = op1(e10, e10)) | ~ sP0), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10))) | ~ sP0), inference(nnf_transformation, [], [f21])).
fof(f1734, plain, (~ spl72_158 | ~ spl72_41 | spl72_46), inference(avatar_split_clause, [], [f550, f1050, f1029, f1729])).
fof(f550, plain, ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11)) | ~ sP0), inference(cnf_transformation, [], [f125])).
fof(f1725, plain, (~ spl72_157 | ~ spl72_22 | spl72_27), inference(avatar_split_clause, [], [f546, f969, f948, f1721])).
fof(f546, plain, ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12)) | ~ sP1), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (op1(e10, e10) = e11)) | ~ sP1), inference(nnf_transformation, [], [f22])).
fof(f1719, plain, (~ spl72_156 | spl72_63), inference(avatar_split_clause, [], [f538, f1122, f1713])).
fof(f538, plain, ((op1(e10, e10) = e12) | ~ sP2), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (op1(e10, e10) = e12)) | ~ sP2), inference(nnf_transformation, [], [f23])).
fof(f1711, plain, (~ spl72_155 | spl72_64), inference(avatar_split_clause, [], [f533, f1126, f1705])).
fof(f533, plain, ((op1(e10, e10) = e13) | ~ sP3), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (op1(e10, e10) = e13)) | ~ sP3), inference(nnf_transformation, [], [f24])).
fof(f1710, plain, (~ spl72_155 | ~ spl72_64 | spl72_49), inference(avatar_split_clause, [], [f534, f1063, f1126, f1705])).
fof(f534, plain, ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13) | ~ sP3), inference(cnf_transformation, [], [f122])).
fof(f1702, plain, (~ spl72_154 | ~ spl72_41 | spl72_46), inference(avatar_split_clause, [], [f530, f1050, f1029, f1697])).
fof(f530, plain, ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11)) | ~ sP4), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e11, e11))) | ~ sP4), inference(nnf_transformation, [], [f25])).
fof(f1700, plain, (~ spl72_154 | ~ spl72_1 | spl72_16), inference(avatar_split_clause, [], [f532, f922, f859, f1697])).
fof(f532, plain, ((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13)) | ~ sP4), inference(cnf_transformation, [], [f121])).
fof(f1693, plain, (~ spl72_153 | ~ spl72_22 | spl72_27), inference(avatar_split_clause, [], [f526, f969, f948, f1689])).
fof(f526, plain, ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12)) | ~ sP5), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e11, e11))) | ~ sP5), inference(nnf_transformation, [], [f26])).
fof(f1687, plain, (~ spl72_152 | spl72_43), inference(avatar_split_clause, [], [f518, f1037, f1681])).
fof(f518, plain, ((e12 = op1(e11, e11)) | ~ sP6), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e11, e11))) | ~ sP6), inference(nnf_transformation, [], [f27])).
fof(f1679, plain, (~ spl72_151 | spl72_44), inference(avatar_split_clause, [], [f513, f1041, f1673])).
fof(f513, plain, ((e13 = op1(e11, e11)) | ~ sP7), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e11, e11))) | ~ sP7), inference(nnf_transformation, [], [f28])).
fof(f1671, plain, (~ spl72_150 | spl72_21), inference(avatar_split_clause, [], [f508, f944, f1665])).
fof(f508, plain, ((e10 = op1(e12, e12)) | ~ sP8), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e12, e12))) | ~ sP8), inference(nnf_transformation, [], [f29])).
fof(f1661, plain, (~ spl72_149 | ~ spl72_22 | spl72_27), inference(avatar_split_clause, [], [f506, f969, f948, f1657])).
fof(f506, plain, ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12)) | ~ sP9), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e12, e12))) | ~ sP9), inference(nnf_transformation, [], [f30])).
fof(f1655, plain, (~ spl72_148 | spl72_23), inference(avatar_split_clause, [], [f498, f952, f1649])).
fof(f498, plain, ((e12 = op1(e12, e12)) | ~ sP10), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e12, e12))) | ~ sP10), inference(nnf_transformation, [], [f31])).
fof(f1647, plain, (~ spl72_147 | spl72_24), inference(avatar_split_clause, [], [f493, f956, f1641])).
fof(f493, plain, ((e13 = op1(e12, e12)) | ~ sP11), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)) & (e13 = op1(e12, e12))) | ~ sP11), inference(nnf_transformation, [], [f32])).
fof(f1639, plain, (~ spl72_146 | spl72_1), inference(avatar_split_clause, [], [f488, f859, f1633])).
fof(f488, plain, ((e10 = op1(e13, e13)) | ~ sP12), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e13, e13))) | ~ sP12), inference(nnf_transformation, [], [f33])).
fof(f1638, plain, (~ spl72_146 | ~ spl72_41 | spl72_46), inference(avatar_split_clause, [], [f490, f1050, f1029, f1633])).
fof(f490, plain, ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11)) | ~ sP12), inference(cnf_transformation, [], [f113])).
fof(f1636, plain, (~ spl72_146 | ~ spl72_1 | spl72_16), inference(avatar_split_clause, [], [f492, f922, f859, f1633])).
fof(f492, plain, ((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13)) | ~ sP12), inference(cnf_transformation, [], [f113])).
fof(f1629, plain, (~ spl72_145 | ~ spl72_22 | spl72_27), inference(avatar_split_clause, [], [f486, f969, f948, f1625])).
fof(f486, plain, ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12)) | ~ sP13), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)) & (e11 = op1(e13, e13))) | ~ sP13), inference(nnf_transformation, [], [f34])).
fof(f1623, plain, (~ spl72_144 | spl72_3), inference(avatar_split_clause, [], [f478, f867, f1617])).
fof(f478, plain, ((e12 = op1(e13, e13)) | ~ sP14), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)) & (e12 = op1(e13, e13))) | ~ sP14), inference(nnf_transformation, [], [f35])).
fof(f1620, plain, (~ spl72_144 | ~ spl72_3 | spl72_8), inference(avatar_split_clause, [], [f482, f888, f867, f1617])).
fof(f482, plain, ((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13)) | ~ sP14), inference(cnf_transformation, [], [f111])).
fof(f1615, plain, (~ spl72_143 | spl72_61), inference(avatar_split_clause, [], [f472, f1114, f1607])).
fof(f472, plain, ((e10 = op1(e10, e10)) | ~ sP15), inference(cnf_transformation, [], [f110])).
fof(f110, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e10 = op1(e10, e10)) & (e10 = op1(e10, e10))) | ~ sP15), inference(nnf_transformation, [], [f36])).
fof(f1613, plain, (~ spl72_143 | ~ spl72_61), inference(avatar_split_clause, [], [f843, f1114, f1607])).
fof(f843, plain, (~ (e10 = op1(e10, e10)) | ~ sP15), inference(duplicate_literal_removal, [], [f474])).
fof(f474, plain, (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)) | ~ sP15), inference(cnf_transformation, [], [f110])).
fof(f1605, plain, (~ spl72_142 | spl72_57), inference(avatar_split_clause, [], [f466, f1097, f1597])).
fof(f466, plain, ((e10 = op1(e10, e11)) | ~ sP16), inference(cnf_transformation, [], [f109])).
fof(f109, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e10 = op1(e11, e10)) & (e10 = op1(e10, e11))) | ~ sP16), inference(nnf_transformation, [], [f37])).
fof(f1595, plain, (~ spl72_141 | spl72_53), inference(avatar_split_clause, [], [f460, f1080, f1587])).
fof(f460, plain, ((e10 = op1(e10, e12)) | ~ sP17), inference(cnf_transformation, [], [f108])).
fof(f108, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e10 = op1(e12, e10)) & (e10 = op1(e10, e12))) | ~ sP17), inference(nnf_transformation, [], [f38])).
fof(f1594, plain, (~ spl72_141 | spl72_29), inference(avatar_split_clause, [], [f461, f978, f1587])).
fof(f461, plain, ((e10 = op1(e12, e10)) | ~ sP17), inference(cnf_transformation, [], [f108])).
fof(f1585, plain, (~ spl72_140 | spl72_49), inference(avatar_split_clause, [], [f454, f1063, f1577])).
fof(f454, plain, ((e10 = op1(e10, e13)) | ~ sP18), inference(cnf_transformation, [], [f107])).
fof(f107, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e10 = op1(e13, e10)) & (e10 = op1(e10, e13))) | ~ sP18), inference(nnf_transformation, [], [f39])).
fof(f1584, plain, (~ spl72_140 | spl72_13), inference(avatar_split_clause, [], [f455, f910, f1577])).
fof(f455, plain, ((e10 = op1(e13, e10)) | ~ sP18), inference(cnf_transformation, [], [f107])).
fof(f1580, plain, (~ spl72_140 | ~ spl72_4), inference(avatar_split_clause, [], [f846, f871, f1577])).
fof(f846, plain, (~ (e13 = op1(e13, e13)) | ~ sP18), inference(duplicate_literal_removal, [], [f459])).
fof(f459, plain, (~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13)) | ~ sP18), inference(cnf_transformation, [], [f107])).
fof(f1575, plain, (~ spl72_139 | spl72_46), inference(avatar_split_clause, [], [f448, f1050, f1567])).
fof(f448, plain, ((e11 = op1(e11, e10)) | ~ sP19), inference(cnf_transformation, [], [f106])).
fof(f106, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e11 = op1(e10, e11)) & (e11 = op1(e11, e10))) | ~ sP19), inference(nnf_transformation, [], [f40])).
fof(f1574, plain, (~ spl72_139 | spl72_58), inference(avatar_split_clause, [], [f449, f1101, f1567])).
fof(f449, plain, ((e11 = op1(e10, e11)) | ~ sP19), inference(cnf_transformation, [], [f106])).
fof(f1573, plain, (~ spl72_139 | ~ spl72_61), inference(avatar_split_clause, [], [f847, f1114, f1567])).
fof(f847, plain, (~ (e10 = op1(e10, e10)) | ~ sP19), inference(duplicate_literal_removal, [], [f450])).
fof(f450, plain, (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)) | ~ sP19), inference(cnf_transformation, [], [f106])).
fof(f1565, plain, (~ spl72_138 | spl72_42), inference(avatar_split_clause, [], [f442, f1033, f1557])).
fof(f1033, plain, (spl72_42 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl72_42])])).
fof(f442, plain, ((e11 = op1(e11, e11)) | ~ sP20), inference(cnf_transformation, [], [f105])).
fof(f105, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e11 = op1(e11, e11)) & (e11 = op1(e11, e11))) | ~ sP20), inference(nnf_transformation, [], [f41])).
fof(f1563, plain, (~ spl72_138 | ~ spl72_46 | ~ spl72_41), inference(avatar_split_clause, [], [f444, f1029, f1050, f1557])).
fof(f444, plain, (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10)) | ~ sP20), inference(cnf_transformation, [], [f105])).
fof(f1562, plain, (~ spl72_138 | ~ spl72_42), inference(avatar_split_clause, [], [f848, f1033, f1557])).
fof(f848, plain, (~ (e11 = op1(e11, e11)) | ~ sP20), inference(duplicate_literal_removal, [], [f445])).
fof(f445, plain, (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11)) | ~ sP20), inference(cnf_transformation, [], [f105])).
fof(f1555, plain, (~ spl72_137 | spl72_38), inference(avatar_split_clause, [], [f436, f1016, f1547])).
fof(f436, plain, ((e11 = op1(e11, e12)) | ~ sP21), inference(cnf_transformation, [], [f104])).
fof(f104, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e11 = op1(e12, e11)) & (e11 = op1(e11, e12))) | ~ sP21), inference(nnf_transformation, [], [f42])).
fof(f1545, plain, (~ spl72_136 | spl72_34), inference(avatar_split_clause, [], [f430, f999, f1537])).
fof(f430, plain, ((e11 = op1(e11, e13)) | ~ sP22), inference(cnf_transformation, [], [f103])).
fof(f103, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e11 = op1(e13, e11)) & (e11 = op1(e11, e13))) | ~ sP22), inference(nnf_transformation, [], [f43])).
fof(f1544, plain, (~ spl72_136 | spl72_10), inference(avatar_split_clause, [], [f431, f897, f1537])).
fof(f431, plain, ((e11 = op1(e13, e11)) | ~ sP22), inference(cnf_transformation, [], [f103])).
fof(f1541, plain, (~ spl72_136 | ~ spl72_8 | ~ spl72_3), inference(avatar_split_clause, [], [f434, f867, f888, f1537])).
fof(f434, plain, (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12)) | ~ sP22), inference(cnf_transformation, [], [f103])).
fof(f1540, plain, (~ spl72_136 | ~ spl72_4), inference(avatar_split_clause, [], [f850, f871, f1537])).
fof(f850, plain, (~ (e13 = op1(e13, e13)) | ~ sP22), inference(duplicate_literal_removal, [], [f435])).
fof(f435, plain, (~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13)) | ~ sP22), inference(cnf_transformation, [], [f103])).
fof(f1535, plain, (~ spl72_135 | spl72_31), inference(avatar_split_clause, [], [f424, f986, f1527])).
fof(f424, plain, ((e12 = op1(e12, e10)) | ~ sP23), inference(cnf_transformation, [], [f102])).
fof(f102, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e12 = op1(e10, e12)) & (e12 = op1(e12, e10))) | ~ sP23), inference(nnf_transformation, [], [f44])).
fof(f1534, plain, (~ spl72_135 | spl72_55), inference(avatar_split_clause, [], [f425, f1088, f1527])).
fof(f425, plain, ((e12 = op1(e10, e12)) | ~ sP23), inference(cnf_transformation, [], [f102])).
fof(f1533, plain, (~ spl72_135 | ~ spl72_61), inference(avatar_split_clause, [], [f851, f1114, f1527])).
fof(f851, plain, (~ (e10 = op1(e10, e10)) | ~ sP23), inference(duplicate_literal_removal, [], [f426])).
fof(f426, plain, (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)) | ~ sP23), inference(cnf_transformation, [], [f102])).
fof(f1525, plain, (~ spl72_134 | spl72_27), inference(avatar_split_clause, [], [f418, f969, f1517])).
fof(f418, plain, ((e12 = op1(e12, e11)) | ~ sP24), inference(cnf_transformation, [], [f101])).
fof(f101, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e12 = op1(e11, e12)) & (e12 = op1(e12, e11))) | ~ sP24), inference(nnf_transformation, [], [f45])).
fof(f1515, plain, (~ spl72_133 | spl72_23), inference(avatar_split_clause, [], [f412, f952, f1507])).
fof(f412, plain, ((e12 = op1(e12, e12)) | ~ sP25), inference(cnf_transformation, [], [f100])).
fof(f100, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e12 = op1(e12, e12)) & (e12 = op1(e12, e12))) | ~ sP25), inference(nnf_transformation, [], [f46])).
fof(f1505, plain, (~ spl72_132 | spl72_19), inference(avatar_split_clause, [], [f406, f935, f1497])).
fof(f406, plain, ((e12 = op1(e12, e13)) | ~ sP26), inference(cnf_transformation, [], [f99])).
fof(f99, plain, (((~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & (~ (e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & (~ (e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & (~ (e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & (e12 = op1(e13, e12)) & (e12 = op1(e12, e13))) | ~ sP26), inference(nnf_transformation, [], [f47])).
fof(f1504, plain, (~ spl72_132 | spl72_7), inference(avatar_split_clause, [], [f407, f884, f1497])).
fof(f407, plain, ((e12 = op1(e13, e12)) | ~ sP26), inference(cnf_transformation, [], [f99])).
fof(f1500, plain, (~ spl72_132 | ~ spl72_4), inference(avatar_split_clause, [], [f854, f871, f1497])).
fof(f854, plain, (~ (e13 = op1(e13, e13)) | ~ sP26), inference(duplicate_literal_removal, [], [f411])).
fof(f411, plain, (~ (e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13)) | ~ sP26), inference(cnf_transformation, [], [f99])).
fof(f1495, plain, (~ spl72_131 | spl72_16), inference(avatar_split_clause, [], [f400, f922, f1487])).
fof(f400, plain, ((e13 = op1(e13, e10)) | ~ sP27), inference(cnf_transformation, [], [f98])).
fof(f98, plain, (((~ (op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & (~ (op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))) & (e13 = op1(e10, e13)) & (e13 = op1(e13, e10))) | ~ sP27), inference(nnf_transformation, [], [f48])).
fof(f1494, plain, (~ spl72_131 | spl72_52), inference(avatar_split_clause, [], [f401, f1075, f1487])).
fof(f401, plain, ((e13 = op1(e10, e13)) | ~ sP27), inference(cnf_transformation, [], [f98])).
fof(f1493, plain, (~ spl72_131 | ~ spl72_61), inference(avatar_split_clause, [], [f855, f1114, f1487])).
fof(f855, plain, (~ (e10 = op1(e10, e10)) | ~ sP27), inference(duplicate_literal_removal, [], [f402])).
fof(f402, plain, (~ (e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)) | ~ sP27), inference(cnf_transformation, [], [f98])).
fof(f1491, plain, (~ spl72_131 | ~ spl72_53 | ~ spl72_63), inference(avatar_split_clause, [], [f404, f1122, f1080, f1487])).
fof(f404, plain, (~ (op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12)) | ~ sP27), inference(cnf_transformation, [], [f98])).
fof(f1485, plain, (~ spl72_130 | spl72_12), inference(avatar_split_clause, [], [f394, f905, f1477])).
fof(f394, plain, ((e13 = op1(e13, e11)) | ~ sP28), inference(cnf_transformation, [], [f97])).
fof(f97, plain, (((~ (e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & (~ (e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & (~ (e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & (~ (e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & (e13 = op1(e11, e13)) & (e13 = op1(e13, e11))) | ~ sP28), inference(nnf_transformation, [], [f49])).
fof(f1474, plain, (~ spl72_129 | spl72_20), inference(avatar_split_clause, [], [f389, f939, f1467])).
fof(f389, plain, ((e13 = op1(e12, e13)) | ~ sP29), inference(cnf_transformation, [], [f96])).
fof(f96, plain, (((~ (e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & (~ (e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & (~ (e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & (~ (e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & (e13 = op1(e12, e13)) & (e13 = op1(e13, e12))) | ~ sP29), inference(nnf_transformation, [], [f50])).
fof(f1465, plain, (spl72_125 | spl72_121 | spl72_117 | spl72_113), inference(avatar_split_clause, [], [f232, f1367, f1384, f1401, f1418])).
fof(f232, plain, ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax4)).
fof(f1464, plain, (spl72_125 | spl72_109 | spl72_93 | spl72_77), inference(avatar_split_clause, [], [f233, f1214, f1282, f1350, f1418])).
fof(f233, plain, ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f1462, plain, (spl72_126 | spl72_110 | spl72_94 | spl72_78), inference(avatar_split_clause, [], [f235, f1218, f1286, f1354, f1422])).
fof(f235, plain, ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)), inference(cnf_transformation, [], [f4])).
fof(f1459, plain, (spl72_128 | spl72_124 | spl72_120 | spl72_116), inference(avatar_split_clause, [], [f238, f1379, f1396, f1413, f1430])).
fof(f238, plain, ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)), inference(cnf_transformation, [], [f4])).
fof(f1458, plain, (spl72_128 | spl72_112 | spl72_96 | spl72_80), inference(avatar_split_clause, [], [f239, f1226, f1294, f1362, f1430])).
fof(f239, plain, ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)), inference(cnf_transformation, [], [f4])).
fof(f1455, plain, (spl72_110 | spl72_106 | spl72_102 | spl72_98), inference(avatar_split_clause, [], [f242, f1303, f1320, f1337, f1354])).
fof(f242, plain, ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1452, plain, (spl72_123 | spl72_107 | spl72_91 | spl72_75), inference(avatar_split_clause, [], [f245, f1205, f1273, f1341, f1409])).
fof(f245, plain, ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1451, plain, (spl72_112 | spl72_108 | spl72_104 | spl72_100), inference(avatar_split_clause, [], [f246, f1311, f1328, f1345, f1362])).
fof(f246, plain, ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f1449, plain, (spl72_93 | spl72_89 | spl72_85 | spl72_81), inference(avatar_split_clause, [], [f248, f1231, f1248, f1265, f1282])).
fof(f248, plain, ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1448, plain, (spl72_117 | spl72_101 | spl72_85 | spl72_69), inference(avatar_split_clause, [], [f249, f1180, f1248, f1316, f1384])).
fof(f249, plain, ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1445, plain, (spl72_95 | spl72_91 | spl72_87 | spl72_83), inference(avatar_split_clause, [], [f252, f1239, f1256, f1273, f1290])).
fof(f252, plain, ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1444, plain, (spl72_119 | spl72_103 | spl72_87 | spl72_71), inference(avatar_split_clause, [], [f253, f1188, f1256, f1324, f1392])).
fof(f253, plain, ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1442, plain, (spl72_120 | spl72_104 | spl72_88 | spl72_72), inference(avatar_split_clause, [], [f255, f1192, f1260, f1328, f1396])).
fof(f255, plain, ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1441, plain, (spl72_77 | spl72_73 | spl72_69 | spl72_65), inference(avatar_split_clause, [], [f256, f1163, f1180, f1197, f1214])).
fof(f256, plain, ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1440, plain, (spl72_113 | spl72_97 | spl72_81 | spl72_65), inference(avatar_split_clause, [], [f257, f1163, f1231, f1299, f1367])).
fof(f257, plain, ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1439, plain, (spl72_78 | spl72_74 | spl72_70 | spl72_66), inference(avatar_split_clause, [], [f258, f1167, f1184, f1201, f1218])).
fof(f258, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1438, plain, (spl72_114 | spl72_98 | spl72_82 | spl72_66), inference(avatar_split_clause, [], [f259, f1167, f1235, f1303, f1371])).
fof(f259, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1435, plain, (spl72_80 | spl72_76 | spl72_72 | spl72_68), inference(avatar_split_clause, [], [f262, f1175, f1192, f1209, f1226])).
fof(f262, plain, ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1434, plain, (spl72_116 | spl72_100 | spl72_84 | spl72_68), inference(avatar_split_clause, [], [f263, f1175, f1243, f1311, f1379])).
fof(f263, plain, ((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1433, plain, (spl72_125 | spl72_126 | spl72_127 | spl72_128), inference(avatar_split_clause, [], [f216, f1430, f1426, f1422, f1418])).
fof(f216, plain, ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax3)).
fof(f1416, plain, (spl72_121 | spl72_122 | spl72_123 | spl72_124), inference(avatar_split_clause, [], [f217, f1413, f1409, f1405, f1401])).
fof(f217, plain, ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))), inference(cnf_transformation, [], [f3])).
fof(f1399, plain, (spl72_117 | spl72_118 | spl72_119 | spl72_120), inference(avatar_split_clause, [], [f218, f1396, f1392, f1388, f1384])).
fof(f218, plain, ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f3])).
fof(f1382, plain, (spl72_113 | spl72_114 | spl72_115 | spl72_116), inference(avatar_split_clause, [], [f219, f1379, f1375, f1371, f1367])).
fof(f219, plain, ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f3])).
fof(f1365, plain, (spl72_109 | spl72_110 | spl72_111 | spl72_112), inference(avatar_split_clause, [], [f220, f1362, f1358, f1354, f1350])).
fof(f220, plain, ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f3])).
fof(f1331, plain, (spl72_101 | spl72_102 | spl72_103 | spl72_104), inference(avatar_split_clause, [], [f222, f1328, f1324, f1320, f1316])).
fof(f222, plain, ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))), inference(cnf_transformation, [], [f3])).
fof(f1314, plain, (spl72_97 | spl72_98 | spl72_99 | spl72_100), inference(avatar_split_clause, [], [f223, f1311, f1307, f1303, f1299])).
fof(f223, plain, ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))), inference(cnf_transformation, [], [f3])).
fof(f1297, plain, (spl72_93 | spl72_94 | spl72_95 | spl72_96), inference(avatar_split_clause, [], [f224, f1294, f1290, f1286, f1282])).
fof(f224, plain, ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f3])).
fof(f1229, plain, (spl72_77 | spl72_78 | spl72_79 | spl72_80), inference(avatar_split_clause, [], [f228, f1226, f1222, f1218, f1214])).
fof(f228, plain, ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))), inference(cnf_transformation, [], [f3])).
fof(f1212, plain, (spl72_73 | spl72_74 | spl72_75 | spl72_76), inference(avatar_split_clause, [], [f229, f1209, f1205, f1201, f1197])).
fof(f229, plain, ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))), inference(cnf_transformation, [], [f3])).
fof(f1195, plain, (spl72_69 | spl72_70 | spl72_71 | spl72_72), inference(avatar_split_clause, [], [f230, f1192, f1188, f1184, f1180])).
fof(f230, plain, ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))), inference(cnf_transformation, [], [f3])).
fof(f1178, plain, (spl72_65 | spl72_66 | spl72_67 | spl72_68), inference(avatar_split_clause, [], [f231, f1175, f1171, f1167, f1163])).
fof(f231, plain, ((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))), inference(cnf_transformation, [], [f3])).
fof(f1161, plain, (spl72_61 | spl72_57 | spl72_53 | spl72_49), inference(avatar_split_clause, [], [f184, f1063, f1080, f1097, f1114])).
fof(f184, plain, ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax2)).
fof(f1160, plain, (spl72_61 | spl72_45 | spl72_29 | spl72_13), inference(avatar_split_clause, [], [f185, f910, f978, f1046, f1114])).
fof(f185, plain, ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f1158, plain, (spl72_62 | spl72_46 | spl72_30 | spl72_14), inference(avatar_split_clause, [], [f187, f914, f982, f1050, f1118])).
fof(f187, plain, ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)), inference(cnf_transformation, [], [f2])).
fof(f1157, plain, (spl72_63 | spl72_59 | spl72_55 | spl72_51), inference(avatar_split_clause, [], [f188, f1071, f1088, f1105, f1122])).
fof(f188, plain, ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)), inference(cnf_transformation, [], [f2])).
fof(f1155, plain, (spl72_64 | spl72_60 | spl72_56 | spl72_52), inference(avatar_split_clause, [], [f190, f1075, f1092, f1109, f1126])).
fof(f190, plain, ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)), inference(cnf_transformation, [], [f2])).
fof(f1154, plain, (spl72_64 | spl72_48 | spl72_32 | spl72_16), inference(avatar_split_clause, [], [f191, f922, f990, f1058, f1126])).
fof(f191, plain, ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)), inference(cnf_transformation, [], [f2])).
fof(f1149, plain, (spl72_47 | spl72_43 | spl72_39 | spl72_35), inference(avatar_split_clause, [], [f196, f1003, f1020, f1037, f1054])).
fof(f196, plain, ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1148, plain, (spl72_59 | spl72_43 | spl72_27 | spl72_11), inference(avatar_split_clause, [], [f197, f901, f969, f1037, f1105])).
fof(f197, plain, ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f1147, plain, (spl72_48 | spl72_44 | spl72_40 | spl72_36), inference(avatar_split_clause, [], [f198, f1007, f1024, f1041, f1058])).
fof(f198, plain, ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f1144, plain, (spl72_53 | spl72_37 | spl72_21 | spl72_5), inference(avatar_split_clause, [], [f201, f876, f944, f1012, f1080])).
fof(f201, plain, ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1141, plain, (spl72_31 | spl72_27 | spl72_23 | spl72_19), inference(avatar_split_clause, [], [f204, f935, f952, f969, f986])).
fof(f204, plain, ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f1140, plain, (spl72_55 | spl72_39 | spl72_23 | spl72_7), inference(avatar_split_clause, [], [f205, f884, f952, f1020, f1088])).
fof(f205, plain, ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1138, plain, (spl72_56 | spl72_40 | spl72_24 | spl72_8), inference(avatar_split_clause, [], [f207, f888, f956, f1024, f1092])).
fof(f207, plain, ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f1136, plain, (spl72_49 | spl72_33 | spl72_17 | spl72_1), inference(avatar_split_clause, [], [f209, f859, f927, f995, f1063])).
fof(f209, plain, ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1135, plain, (spl72_14 | spl72_10 | spl72_6 | spl72_2), inference(avatar_split_clause, [], [f210, f863, f880, f897, f914])).
fof(f210, plain, ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1134, plain, (spl72_50 | spl72_34 | spl72_18 | spl72_2), inference(avatar_split_clause, [], [f211, f863, f931, f999, f1067])).
fof(f211, plain, ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1132, plain, (spl72_51 | spl72_35 | spl72_19 | spl72_3), inference(avatar_split_clause, [], [f213, f867, f935, f1003, f1071])).
fof(f213, plain, ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1131, plain, (spl72_16 | spl72_12 | spl72_8 | spl72_4), inference(avatar_split_clause, [], [f214, f871, f888, f905, f922])).
fof(f214, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f1130, plain, (spl72_52 | spl72_36 | spl72_20 | spl72_4), inference(avatar_split_clause, [], [f215, f871, f939, f1007, f1075])).
fof(f215, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f1129, plain, (spl72_61 | spl72_62 | spl72_63 | spl72_64), inference(avatar_split_clause, [], [f168, f1126, f1122, f1118, f1114])).
fof(f168, plain, ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG128+1.p', ax1)).
fof(f1112, plain, (spl72_57 | spl72_58 | spl72_59 | spl72_60), inference(avatar_split_clause, [], [f169, f1109, f1105, f1101, f1097])).
fof(f169, plain, ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f1])).
fof(f1095, plain, (spl72_53 | spl72_54 | spl72_55 | spl72_56), inference(avatar_split_clause, [], [f170, f1092, f1088, f1084, f1080])).
fof(f170, plain, ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f1])).
fof(f1078, plain, (spl72_49 | spl72_50 | spl72_51 | spl72_52), inference(avatar_split_clause, [], [f171, f1075, f1071, f1067, f1063])).
fof(f171, plain, ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f1])).
fof(f1061, plain, (spl72_45 | spl72_46 | spl72_47 | spl72_48), inference(avatar_split_clause, [], [f172, f1058, f1054, f1050, f1046])).
fof(f172, plain, ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f1])).
fof(f1027, plain, (spl72_37 | spl72_38 | spl72_39 | spl72_40), inference(avatar_split_clause, [], [f174, f1024, f1020, f1016, f1012])).
fof(f174, plain, ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))), inference(cnf_transformation, [], [f1])).
fof(f1010, plain, (spl72_33 | spl72_34 | spl72_35 | spl72_36), inference(avatar_split_clause, [], [f175, f1007, f1003, f999, f995])).
fof(f175, plain, ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))), inference(cnf_transformation, [], [f1])).
fof(f993, plain, (spl72_29 | spl72_30 | spl72_31 | spl72_32), inference(avatar_split_clause, [], [f176, f990, f986, f982, f978])).
fof(f176, plain, ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f1])).
fof(f942, plain, (spl72_17 | spl72_18 | spl72_19 | spl72_20), inference(avatar_split_clause, [], [f179, f939, f935, f931, f927])).
fof(f179, plain, ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))), inference(cnf_transformation, [], [f1])).
fof(f925, plain, (spl72_13 | spl72_14 | spl72_15 | spl72_16), inference(avatar_split_clause, [], [f180, f922, f918, f914, f910])).
fof(f180, plain, ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f1])).
fof(f891, plain, (spl72_5 | spl72_6 | spl72_7 | spl72_8), inference(avatar_split_clause, [], [f182, f888, f884, f880, f876])).
fof(f182, plain, ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))), inference(cnf_transformation, [], [f1])).
fof(f874, plain, (spl72_1 | spl72_2 | spl72_3 | spl72_4), inference(avatar_split_clause, [], [f183, f871, f867, f863, f859])).
fof(f183, plain, ((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))), inference(cnf_transformation, [], [f1])).