fof(f5096, plain, $false, inference(avatar_sat_refutation, [], [f641, f692, f726, f777, f811, f828, f858, f859, f868, f869, f1013, f1081, f1115, f1132, f1152, f1153, f1155, f1159, f1172, f1173, f1182, f1183, f1184, f1185, f1186, f1187, f1188, f1189, f1190, f1191, f1547, f3031, f3064, f3068, f3080, f3083, f3108, f3112, f3124, f3146, f3180, f3202, f3204, f3205, f3206, f3255, f3257, f3258, f3268, f3299, f3305, f3336, f3344, f3346, f3358, f3360, f3364, f3369, f3377, f3381, f3385, f3407, f3411, f3414, f3417, f3432, f3436, f3440, f3454, f3470, f3477, f3499, f3528, f3549, f3565, f3570, f3584, f3600, f3614, f3630, f3637, f3656, f3671, f3674, f3711, f3718, f3722, f3725, f3734, f3766, f3807, f3809, f3893, f3968, f4072, f4172, f4264, f4351, f4436, f4513, f4602, f4682, f4689, f4761, f4841, f4921, f5015])).
fof(f5015, plain, (~ spl36_48 | ~ spl36_70 | ~ spl36_178 | spl36_378 | ~ spl36_380), inference(avatar_contradiction_clause, [], [f5014])).
fof(f5014, plain, ($false | (~ spl36_48 | ~ spl36_70 | ~ spl36_178 | spl36_378 | ~ spl36_380)), inference(subsumption_resolution, [], [f5013, f902])).
fof(f902, plain, ((e21 = op2(e23, e22)) | ~ spl36_70), inference(avatar_component_clause, [], [f900])).
fof(f900, plain, (spl36_70 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl36_70])])).
fof(f5013, plain, (~ (e21 = op2(e23, e22)) | (~ spl36_48 | ~ spl36_178 | spl36_378 | ~ spl36_380)), inference(forward_demodulation, [], [f5012, f335])).
fof(f335, plain, (e21 = h1(e13)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((h1(e11) = op2(op2(e20, e21), e21)) & (op2(e20, e21) = h1(e10)) & (e21 = h1(e13)) & (e20 = h1(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax14)).
fof(f5012, plain, (~ (op2(e23, e22) = h1(e13)) | (~ spl36_48 | ~ spl36_178 | spl36_378 | ~ spl36_380)), inference(forward_demodulation, [], [f5011, f776])).
fof(f776, plain, ((e13 = op1(e11, e10)) | ~ spl36_48), inference(avatar_component_clause, [], [f774])).
fof(f774, plain, (spl36_48 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_48])])).
fof(f5011, plain, (~ (op2(e23, e22) = h1(op1(e11, e10))) | (~ spl36_178 | spl36_378 | ~ spl36_380)), inference(forward_demodulation, [], [f5010, f2981])).
fof(f2981, plain, ((e23 = h1(e11)) | ~ spl36_380), inference(avatar_component_clause, [], [f2980])).
fof(f2980, plain, (spl36_380 <=> (e23 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_380])])).
fof(f5010, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), e22)) | (~ spl36_178 | spl36_378)), inference(forward_demodulation, [], [f2974, f1545])).
fof(f1545, plain, ((e22 = h1(e10)) | ~ spl36_178), inference(avatar_component_clause, [], [f1544])).
fof(f1544, plain, (spl36_178 <=> (e22 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_178])])).
fof(f2974, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | spl36_378), inference(avatar_component_clause, [], [f2972])).
fof(f2972, plain, (spl36_378 <=> (h1(op1(e11, e10)) = op2(h1(e11), h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl36_378])])).
fof(f4921, plain, (~ spl36_6 | ~ spl36_339 | ~ spl36_380 | spl36_382), inference(avatar_contradiction_clause, [], [f4920])).
fof(f4920, plain, ($false | (~ spl36_6 | ~ spl36_339 | ~ spl36_380 | spl36_382)), inference(subsumption_resolution, [], [f4919, f2981])).
fof(f4919, plain, (~ (e23 = h1(e11)) | (~ spl36_6 | ~ spl36_339 | spl36_382)), inference(forward_demodulation, [], [f4918, f2685])).
fof(f2685, plain, ((e23 = h4(e10)) | ~ spl36_339), inference(avatar_component_clause, [], [f2684])).
fof(f2684, plain, (spl36_339 <=> (e23 = h4(e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_339])])).
fof(f4918, plain, (~ (h1(e11) = h4(e10)) | (~ spl36_6 | spl36_382)), inference(forward_demodulation, [], [f2990, f598])).
fof(f598, plain, ((e11 = op1(e13, e12)) | ~ spl36_6), inference(avatar_component_clause, [], [f596])).
fof(f596, plain, (spl36_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl36_6])])).
fof(f2990, plain, (~ (h4(e10) = h1(op1(e13, e12))) | spl36_382), inference(avatar_component_clause, [], [f2988])).
fof(f2988, plain, (spl36_382 <=> (h4(e10) = h1(op1(e13, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl36_382])])).
fof(f4841, plain, (~ spl36_59 | ~ spl36_81 | ~ spl36_178 | spl36_377 | ~ spl36_380), inference(avatar_contradiction_clause, [], [f4840])).
fof(f4840, plain, ($false | (~ spl36_59 | ~ spl36_81 | ~ spl36_178 | spl36_377 | ~ spl36_380)), inference(subsumption_resolution, [], [f4839, f949])).
fof(f949, plain, ((e20 = op2(e22, e23)) | ~ spl36_81), inference(avatar_component_clause, [], [f947])).
fof(f947, plain, (spl36_81 <=> (e20 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl36_81])])).
fof(f4839, plain, (~ (e20 = op2(e22, e23)) | (~ spl36_59 | ~ spl36_178 | spl36_377 | ~ spl36_380)), inference(forward_demodulation, [], [f4838, f334])).
fof(f334, plain, (e20 = h1(e12)), inference(cnf_transformation, [], [f14])).
fof(f4838, plain, (~ (op2(e22, e23) = h1(e12)) | (~ spl36_59 | ~ spl36_178 | spl36_377 | ~ spl36_380)), inference(forward_demodulation, [], [f4837, f823])).
fof(f823, plain, ((e12 = op1(e10, e11)) | ~ spl36_59), inference(avatar_component_clause, [], [f821])).
fof(f821, plain, (spl36_59 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_59])])).
fof(f4837, plain, (~ (op2(e22, e23) = h1(op1(e10, e11))) | (~ spl36_178 | spl36_377 | ~ spl36_380)), inference(forward_demodulation, [], [f4836, f1545])).
fof(f4836, plain, (~ (h1(op1(e10, e11)) = op2(h1(e10), e23)) | (spl36_377 | ~ spl36_380)), inference(forward_demodulation, [], [f2970, f2981])).
fof(f2970, plain, (~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | spl36_377), inference(avatar_component_clause, [], [f2968])).
fof(f2968, plain, (spl36_377 <=> (h1(op1(e10, e11)) = op2(h1(e10), h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl36_377])])).
fof(f4761, plain, (~ spl36_42 | ~ spl36_68 | spl36_379 | ~ spl36_380), inference(avatar_contradiction_clause, [], [f4760])).
fof(f4760, plain, ($false | (~ spl36_42 | ~ spl36_68 | spl36_379 | ~ spl36_380)), inference(subsumption_resolution, [], [f4759, f893])).
fof(f893, plain, ((e23 = op2(e23, e23)) | ~ spl36_68), inference(avatar_component_clause, [], [f891])).
fof(f891, plain, (spl36_68 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl36_68])])).
fof(f4759, plain, (~ (e23 = op2(e23, e23)) | (~ spl36_42 | spl36_379 | ~ spl36_380)), inference(forward_demodulation, [], [f4758, f2981])).
fof(f4758, plain, (~ (op2(e23, e23) = h1(e11)) | (~ spl36_42 | spl36_379 | ~ spl36_380)), inference(forward_demodulation, [], [f4757, f751])).
fof(f751, plain, ((e11 = op1(e11, e11)) | ~ spl36_42), inference(avatar_component_clause, [], [f749])).
fof(f749, plain, (spl36_42 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_42])])).
fof(f4757, plain, (~ (op2(e23, e23) = h1(op1(e11, e11))) | (spl36_379 | ~ spl36_380)), inference(forward_demodulation, [], [f2978, f2981])).
fof(f2978, plain, (~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | spl36_379), inference(avatar_component_clause, [], [f2976])).
fof(f2976, plain, (spl36_379 <=> (h1(op1(e11, e11)) = op2(h1(e11), h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl36_379])])).
fof(f4689, plain, (~ spl36_4 | ~ spl36_106 | spl36_381), inference(avatar_contradiction_clause, [], [f4688])).
fof(f4688, plain, ($false | (~ spl36_4 | ~ spl36_106 | spl36_381)), inference(subsumption_resolution, [], [f4687, f1055])).
fof(f1055, plain, ((e21 = op2(e21, e21)) | ~ spl36_106), inference(avatar_component_clause, [], [f1053])).
fof(f1053, plain, (spl36_106 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_106])])).
fof(f4687, plain, (~ (e21 = op2(e21, e21)) | (~ spl36_4 | spl36_381)), inference(forward_demodulation, [], [f4686, f335])).
fof(f4686, plain, (~ (op2(e21, e21) = h1(e13)) | (~ spl36_4 | spl36_381)), inference(forward_demodulation, [], [f2986, f589])).
fof(f589, plain, ((e13 = op1(e13, e13)) | ~ spl36_4), inference(avatar_component_clause, [], [f587])).
fof(f587, plain, (spl36_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl36_4])])).
fof(f2986, plain, (~ (op2(e21, e21) = h1(op1(e13, e13))) | spl36_381), inference(avatar_component_clause, [], [f2984])).
fof(f2984, plain, (spl36_381 <=> (op2(e21, e21) = h1(op1(e13, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl36_381])])).
fof(f4682, plain, (spl36_50 | ~ spl36_17), inference(avatar_split_clause, [], [f3340, f643, f783])).
fof(f783, plain, (spl36_50 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl36_50])])).
fof(f643, plain, (spl36_17 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl36_17])])).
fof(f3340, plain, ((e11 = op1(e10, e13)) | ~ spl36_17), inference(forward_demodulation, [], [f331, f645])).
fof(f645, plain, ((e10 = op1(e12, e13)) | ~ spl36_17), inference(avatar_component_clause, [], [f643])).
fof(f331, plain, (e11 = op1(op1(e12, e13), e13)), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((e11 = op1(op1(e12, e13), e13)) & (e10 = op1(e12, e13))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax12)).
fof(f4602, plain, (~ spl36_15 | ~ spl36_101 | ~ spl36_178 | spl36_384), inference(avatar_contradiction_clause, [], [f4601])).
fof(f4601, plain, ($false | (~ spl36_15 | ~ spl36_101 | ~ spl36_178 | spl36_384)), inference(subsumption_resolution, [], [f4600, f1034])).
fof(f1034, plain, ((e20 = op2(e21, e22)) | ~ spl36_101), inference(avatar_component_clause, [], [f1032])).
fof(f1032, plain, (spl36_101 <=> (e20 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl36_101])])).
fof(f4600, plain, (~ (e20 = op2(e21, e22)) | (~ spl36_15 | ~ spl36_178 | spl36_384)), inference(forward_demodulation, [], [f4599, f334])).
fof(f4599, plain, (~ (op2(e21, e22) = h1(e12)) | (~ spl36_15 | ~ spl36_178 | spl36_384)), inference(forward_demodulation, [], [f4598, f636])).
fof(f636, plain, ((e12 = op1(e13, e10)) | ~ spl36_15), inference(avatar_component_clause, [], [f634])).
fof(f634, plain, (spl36_15 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_15])])).
fof(f4598, plain, (~ (op2(e21, e22) = h1(op1(e13, e10))) | (~ spl36_178 | spl36_384)), inference(forward_demodulation, [], [f2998, f1545])).
fof(f2998, plain, (~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | spl36_384), inference(avatar_component_clause, [], [f2996])).
fof(f2996, plain, (spl36_384 <=> (h1(op1(e13, e10)) = op2(e21, h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl36_384])])).
fof(f4513, plain, (~ spl36_17 | spl36_385), inference(avatar_contradiction_clause, [], [f4512])).
fof(f4512, plain, ($false | (~ spl36_17 | spl36_385)), inference(trivial_inequality_removal, [], [f4511])).
fof(f4511, plain, (~ (h1(e10) = h1(e10)) | (~ spl36_17 | spl36_385)), inference(forward_demodulation, [], [f3002, f645])).
fof(f3002, plain, (~ (h1(e10) = h1(op1(e12, e13))) | spl36_385), inference(avatar_component_clause, [], [f3000])).
fof(f3000, plain, (spl36_385 <=> (h1(e10) = h1(op1(e12, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl36_385])])).
fof(f4436, plain, (~ spl36_9 | ~ spl36_99 | ~ spl36_178 | ~ spl36_380 | spl36_383), inference(avatar_contradiction_clause, [], [f4435])).
fof(f4435, plain, ($false | (~ spl36_9 | ~ spl36_99 | ~ spl36_178 | ~ spl36_380 | spl36_383)), inference(subsumption_resolution, [], [f4434, f1025])).
fof(f1025, plain, ((e22 = op2(e21, e23)) | ~ spl36_99), inference(avatar_component_clause, [], [f1023])).
fof(f1023, plain, (spl36_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl36_99])])).
fof(f4434, plain, (~ (e22 = op2(e21, e23)) | (~ spl36_9 | ~ spl36_178 | ~ spl36_380 | spl36_383)), inference(forward_demodulation, [], [f4433, f1545])).
fof(f4433, plain, (~ (op2(e21, e23) = h1(e10)) | (~ spl36_9 | ~ spl36_380 | spl36_383)), inference(forward_demodulation, [], [f4432, f611])).
fof(f611, plain, ((e10 = op1(e13, e11)) | ~ spl36_9), inference(avatar_component_clause, [], [f609])).
fof(f609, plain, (spl36_9 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_9])])).
fof(f4432, plain, (~ (op2(e21, e23) = h1(op1(e13, e11))) | (~ spl36_380 | spl36_383)), inference(forward_demodulation, [], [f2994, f2981])).
fof(f2994, plain, (~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | spl36_383), inference(avatar_component_clause, [], [f2992])).
fof(f2992, plain, (spl36_383 <=> (h1(op1(e13, e11)) = op2(e21, h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl36_383])])).
fof(f4351, plain, (~ spl36_35 | ~ spl36_73 | ~ spl36_380 | spl36_389), inference(avatar_contradiction_clause, [], [f4350])).
fof(f4350, plain, ($false | (~ spl36_35 | ~ spl36_73 | ~ spl36_380 | spl36_389)), inference(subsumption_resolution, [], [f4349, f915])).
fof(f915, plain, ((e20 = op2(e23, e21)) | ~ spl36_73), inference(avatar_component_clause, [], [f913])).
fof(f913, plain, (spl36_73 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_73])])).
fof(f4349, plain, (~ (e20 = op2(e23, e21)) | (~ spl36_35 | ~ spl36_380 | spl36_389)), inference(forward_demodulation, [], [f4348, f334])).
fof(f4348, plain, (~ (op2(e23, e21) = h1(e12)) | (~ spl36_35 | ~ spl36_380 | spl36_389)), inference(forward_demodulation, [], [f4347, f721])).
fof(f721, plain, ((e12 = op1(e11, e13)) | ~ spl36_35), inference(avatar_component_clause, [], [f719])).
fof(f719, plain, (spl36_35 <=> (e12 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl36_35])])).
fof(f4347, plain, (~ (op2(e23, e21) = h1(op1(e11, e13))) | (~ spl36_380 | spl36_389)), inference(forward_demodulation, [], [f3018, f2981])).
fof(f3018, plain, (~ (h1(op1(e11, e13)) = op2(h1(e11), e21)) | spl36_389), inference(avatar_component_clause, [], [f3016])).
fof(f3016, plain, (spl36_389 <=> (h1(op1(e11, e13)) = op2(h1(e11), e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_389])])).
fof(f4264, plain, (~ spl36_28 | ~ spl36_114 | ~ spl36_380 | spl36_387), inference(avatar_contradiction_clause, [], [f4263])).
fof(f4263, plain, ($false | (~ spl36_28 | ~ spl36_114 | ~ spl36_380 | spl36_387)), inference(subsumption_resolution, [], [f4262, f1089])).
fof(f1089, plain, ((e21 = op2(e20, e23)) | ~ spl36_114), inference(avatar_component_clause, [], [f1087])).
fof(f1087, plain, (spl36_114 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl36_114])])).
fof(f4262, plain, (~ (e21 = op2(e20, e23)) | (~ spl36_28 | ~ spl36_380 | spl36_387)), inference(forward_demodulation, [], [f4261, f335])).
fof(f4261, plain, (~ (op2(e20, e23) = h1(e13)) | (~ spl36_28 | ~ spl36_380 | spl36_387)), inference(forward_demodulation, [], [f4260, f691])).
fof(f691, plain, ((e13 = op1(e12, e11)) | ~ spl36_28), inference(avatar_component_clause, [], [f689])).
fof(f689, plain, (spl36_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_28])])).
fof(f4260, plain, (~ (op2(e20, e23) = h1(op1(e12, e11))) | (~ spl36_380 | spl36_387)), inference(forward_demodulation, [], [f3010, f2981])).
fof(f3010, plain, (~ (h1(op1(e12, e11)) = op2(e20, h1(e11))) | spl36_387), inference(avatar_component_clause, [], [f3008])).
fof(f3008, plain, (spl36_387 <=> (h1(op1(e12, e11)) = op2(e20, h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl36_387])])).
fof(f4172, plain, (~ spl36_23 | ~ spl36_125 | spl36_386), inference(avatar_contradiction_clause, [], [f4171])).
fof(f4171, plain, ($false | (~ spl36_23 | ~ spl36_125 | spl36_386)), inference(subsumption_resolution, [], [f4170, f1136])).
fof(f1136, plain, ((e20 = op2(e20, e20)) | ~ spl36_125), inference(avatar_component_clause, [], [f1134])).
fof(f1134, plain, (spl36_125 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_125])])).
fof(f4170, plain, (~ (e20 = op2(e20, e20)) | (~ spl36_23 | spl36_386)), inference(forward_demodulation, [], [f4169, f334])).
fof(f4169, plain, (~ (op2(e20, e20) = h1(e12)) | (~ spl36_23 | spl36_386)), inference(forward_demodulation, [], [f3006, f670])).
fof(f670, plain, ((e12 = op1(e12, e12)) | ~ spl36_23), inference(avatar_component_clause, [], [f668])).
fof(f668, plain, (spl36_23 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl36_23])])).
fof(f3006, plain, (~ (op2(e20, e20) = h1(op1(e12, e12))) | spl36_386), inference(avatar_component_clause, [], [f3004])).
fof(f3004, plain, (spl36_386 <=> (op2(e20, e20) = h1(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl36_386])])).
fof(f4072, plain, (~ spl36_37 | ~ spl36_79 | ~ spl36_178 | ~ spl36_380 | spl36_390), inference(avatar_contradiction_clause, [], [f4071])).
fof(f4071, plain, ($false | (~ spl36_37 | ~ spl36_79 | ~ spl36_178 | ~ spl36_380 | spl36_390)), inference(subsumption_resolution, [], [f4070, f940])).
fof(f940, plain, ((e22 = op2(e23, e20)) | ~ spl36_79), inference(avatar_component_clause, [], [f938])).
fof(f938, plain, (spl36_79 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_79])])).
fof(f4070, plain, (~ (e22 = op2(e23, e20)) | (~ spl36_37 | ~ spl36_178 | ~ spl36_380 | spl36_390)), inference(forward_demodulation, [], [f4069, f1545])).
fof(f4069, plain, (~ (op2(e23, e20) = h1(e10)) | (~ spl36_37 | ~ spl36_380 | spl36_390)), inference(forward_demodulation, [], [f4068, f730])).
fof(f730, plain, ((e10 = op1(e11, e12)) | ~ spl36_37), inference(avatar_component_clause, [], [f728])).
fof(f728, plain, (spl36_37 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl36_37])])).
fof(f4068, plain, (~ (op2(e23, e20) = h1(op1(e11, e12))) | (~ spl36_380 | spl36_390)), inference(forward_demodulation, [], [f3022, f2981])).
fof(f3022, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | spl36_390), inference(avatar_component_clause, [], [f3020])).
fof(f3020, plain, (spl36_390 <=> (h1(op1(e11, e12)) = op2(h1(e11), e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_390])])).
fof(f3968, plain, (~ spl36_56 | ~ spl36_94 | ~ spl36_178 | spl36_392), inference(avatar_contradiction_clause, [], [f3967])).
fof(f3967, plain, ($false | (~ spl36_56 | ~ spl36_94 | ~ spl36_178 | spl36_392)), inference(subsumption_resolution, [], [f3966, f1004])).
fof(f1004, plain, ((e21 = op2(e22, e20)) | ~ spl36_94), inference(avatar_component_clause, [], [f1002])).
fof(f1002, plain, (spl36_94 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_94])])).
fof(f3966, plain, (~ (e21 = op2(e22, e20)) | (~ spl36_56 | ~ spl36_178 | spl36_392)), inference(forward_demodulation, [], [f3965, f335])).
fof(f3965, plain, (~ (op2(e22, e20) = h1(e13)) | (~ spl36_56 | ~ spl36_178 | spl36_392)), inference(forward_demodulation, [], [f3964, f810])).
fof(f810, plain, ((e13 = op1(e10, e12)) | ~ spl36_56), inference(avatar_component_clause, [], [f808])).
fof(f808, plain, (spl36_56 <=> (e13 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl36_56])])).
fof(f3964, plain, (~ (op2(e22, e20) = h1(op1(e10, e12))) | (~ spl36_178 | spl36_392)), inference(forward_demodulation, [], [f3030, f1545])).
fof(f3030, plain, (~ (h1(op1(e10, e12)) = op2(h1(e10), e20)) | spl36_392), inference(avatar_component_clause, [], [f3028])).
fof(f3028, plain, (spl36_392 <=> (h1(op1(e10, e12)) = op2(h1(e10), e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_392])])).
fof(f3893, plain, (~ spl36_380 | ~ spl36_30 | ~ spl36_120 | ~ spl36_178 | spl36_388), inference(avatar_split_clause, [], [f3892, f3012, f1544, f1112, f698, f2980])).
fof(f698, plain, (spl36_30 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_30])])).
fof(f1112, plain, (spl36_120 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl36_120])])).
fof(f3012, plain, (spl36_388 <=> (h1(op1(e12, e10)) = op2(e20, h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl36_388])])).
fof(f3892, plain, (~ (e23 = h1(e11)) | (~ spl36_30 | ~ spl36_120 | ~ spl36_178 | spl36_388)), inference(forward_demodulation, [], [f3891, f1114])).
fof(f1114, plain, ((e23 = op2(e20, e22)) | ~ spl36_120), inference(avatar_component_clause, [], [f1112])).
fof(f3891, plain, (~ (op2(e20, e22) = h1(e11)) | (~ spl36_30 | ~ spl36_178 | spl36_388)), inference(forward_demodulation, [], [f3890, f700])).
fof(f700, plain, ((e11 = op1(e12, e10)) | ~ spl36_30), inference(avatar_component_clause, [], [f698])).
fof(f3890, plain, (~ (op2(e20, e22) = h1(op1(e12, e10))) | (~ spl36_178 | spl36_388)), inference(forward_demodulation, [], [f3014, f1545])).
fof(f3014, plain, (~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | spl36_388), inference(avatar_component_clause, [], [f3012])).
fof(f3809, plain, (~ spl36_61 | ~ spl36_87 | ~ spl36_178 | spl36_376), inference(avatar_contradiction_clause, [], [f3808])).
fof(f3808, plain, ($false | (~ spl36_61 | ~ spl36_87 | ~ spl36_178 | spl36_376)), inference(subsumption_resolution, [], [f3805, f974])).
fof(f974, plain, ((e22 = op2(e22, e22)) | ~ spl36_87), inference(avatar_component_clause, [], [f972])).
fof(f972, plain, (spl36_87 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl36_87])])).
fof(f3805, plain, (~ (e22 = op2(e22, e22)) | (~ spl36_61 | ~ spl36_178 | spl36_376)), inference(backward_demodulation, [], [f3554, f1545])).
fof(f3554, plain, (~ (h1(e10) = op2(h1(e10), h1(e10))) | (~ spl36_61 | spl36_376)), inference(forward_demodulation, [], [f2966, f832])).
fof(f832, plain, ((e10 = op1(e10, e10)) | ~ spl36_61), inference(avatar_component_clause, [], [f830])).
fof(f830, plain, (spl36_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_61])])).
fof(f2966, plain, (~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))) | spl36_376), inference(avatar_component_clause, [], [f2964])).
fof(f2964, plain, (spl36_376 <=> (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), introduced(avatar_definition, [new_symbols(naming, [spl36_376])])).
fof(f3807, plain, (spl36_380 | ~ spl36_92 | ~ spl36_178), inference(avatar_split_clause, [], [f3806, f1544, f993, f2980])).
fof(f993, plain, (spl36_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_92])])).
fof(f3806, plain, ((e23 = h1(e11)) | (~ spl36_92 | ~ spl36_178)), inference(forward_demodulation, [], [f3803, f995])).
fof(f995, plain, ((e23 = op2(e22, e21)) | ~ spl36_92), inference(avatar_component_clause, [], [f993])).
fof(f3803, plain, ((op2(e22, e21) = h1(e11)) | ~ spl36_178), inference(backward_demodulation, [], [f1198, f1545])).
fof(f1198, plain, (h1(e11) = op2(h1(e10), e21)), inference(backward_demodulation, [], [f337, f336])).
fof(f336, plain, (op2(e20, e21) = h1(e10)), inference(cnf_transformation, [], [f14])).
fof(f337, plain, (h1(e11) = op2(op2(e20, e21), e21)), inference(cnf_transformation, [], [f14])).
fof(f3766, plain, (~ spl36_71 | ~ spl36_87), inference(avatar_contradiction_clause, [], [f3765])).
fof(f3765, plain, ($false | (~ spl36_71 | ~ spl36_87)), inference(subsumption_resolution, [], [f3113, f3755])).
fof(f3755, plain, ((e22 = h12(e10)) | ~ spl36_71), inference(forward_demodulation, [], [f380, f906])).
fof(f906, plain, ((e22 = op2(e23, e22)) | ~ spl36_71), inference(avatar_component_clause, [], [f904])).
fof(f904, plain, (spl36_71 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl36_71])])).
fof(f380, plain, (op2(e23, e22) = h12(e10)), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ((h12(e11) = op2(op2(e23, e22), e22)) & (op2(e23, e22) = h12(e10)) & (e22 = h12(e13)) & (e23 = h12(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax25)).
fof(f3113, plain, (~ (e22 = h12(e10)) | ~ spl36_87), inference(backward_demodulation, [], [f1250, f974])).
fof(f1250, plain, ~ (op2(e22, e22) = h12(e10)), inference(backward_demodulation, [], [f263, f380])).
fof(f263, plain, ~ (op2(e22, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax6)).
fof(f3734, plain, (~ spl36_81 | ~ spl36_83), inference(avatar_contradiction_clause, [], [f3733])).
fof(f3733, plain, ($false | (~ spl36_81 | ~ spl36_83)), inference(subsumption_resolution, [], [f3732, f301])).
fof(f301, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax8)).
fof(f3732, plain, ((e20 = e22) | (~ spl36_81 | ~ spl36_83)), inference(forward_demodulation, [], [f957, f949])).
fof(f957, plain, ((e22 = op2(e22, e23)) | ~ spl36_83), inference(avatar_component_clause, [], [f955])).
fof(f955, plain, (spl36_83 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl36_83])])).
fof(f3725, plain, (~ spl36_105 | ~ spl36_106), inference(avatar_contradiction_clause, [], [f3724])).
fof(f3724, plain, ($false | (~ spl36_105 | ~ spl36_106)), inference(subsumption_resolution, [], [f3723, f300])).
fof(f300, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f8])).
fof(f3723, plain, ((e20 = e21) | (~ spl36_105 | ~ spl36_106)), inference(backward_demodulation, [], [f1055, f1051])).
fof(f1051, plain, ((e20 = op2(e21, e21)) | ~ spl36_105), inference(avatar_component_clause, [], [f1049])).
fof(f1049, plain, (spl36_105 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_105])])).
fof(f3722, plain, (spl36_166 | ~ spl36_111), inference(avatar_split_clause, [], [f3721, f1074, f1472])).
fof(f1472, plain, (spl36_166 <=> (e22 = h4(e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_166])])).
fof(f1074, plain, (spl36_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_111])])).
fof(f3721, plain, ((e22 = h4(e10)) | ~ spl36_111), inference(backward_demodulation, [], [f348, f1076])).
fof(f1076, plain, ((e22 = op2(e21, e20)) | ~ spl36_111), inference(avatar_component_clause, [], [f1074])).
fof(f348, plain, (op2(e21, e20) = h4(e10)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((h4(e11) = op2(op2(e21, e20), e20)) & (op2(e21, e20) = h4(e10)) & (e20 = h4(e13)) & (e21 = h4(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax17)).
fof(f3718, plain, (~ spl36_175 | ~ spl36_114), inference(avatar_split_clause, [], [f3714, f1087, f1526])).
fof(f1526, plain, (spl36_175 <=> (e21 = h2(e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_175])])).
fof(f3714, plain, (~ (e21 = h2(e10)) | ~ spl36_114), inference(backward_demodulation, [], [f1203, f1089])).
fof(f1203, plain, ~ (op2(e20, e23) = h2(e10)), inference(backward_demodulation, [], [f275, f340])).
fof(f340, plain, (op2(e20, e22) = h2(e10)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((h2(e11) = op2(op2(e20, e22), e22)) & (op2(e20, e22) = h2(e10)) & (e22 = h2(e13)) & (e20 = h2(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax15)).
fof(f275, plain, ~ (op2(e20, e22) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f3711, plain, (spl36_178 | ~ spl36_123), inference(avatar_split_clause, [], [f3710, f1125, f1544])).
fof(f1125, plain, (spl36_123 <=> (e22 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_123])])).
fof(f3710, plain, ((e22 = h1(e10)) | ~ spl36_123), inference(backward_demodulation, [], [f336, f1127])).
fof(f1127, plain, ((e22 = op2(e20, e21)) | ~ spl36_123), inference(avatar_component_clause, [], [f1125])).
fof(f3674, plain, (spl36_175 | ~ spl36_118), inference(avatar_split_clause, [], [f3673, f1104, f1526])).
fof(f1104, plain, (spl36_118 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl36_118])])).
fof(f3673, plain, ((e21 = h2(e10)) | ~ spl36_118), inference(forward_demodulation, [], [f340, f1106])).
fof(f1106, plain, ((e21 = op2(e20, e22)) | ~ spl36_118), inference(avatar_component_clause, [], [f1104])).
fof(f3671, plain, (spl36_114 | ~ spl36_148), inference(avatar_split_clause, [], [f3558, f1365, f1087])).
fof(f1365, plain, (spl36_148 <=> (e20 = h9(e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_148])])).
fof(f3558, plain, ((e21 = op2(e20, e23)) | ~ spl36_148), inference(backward_demodulation, [], [f1240, f1366])).
fof(f1366, plain, ((e20 = h9(e10)) | ~ spl36_148), inference(avatar_component_clause, [], [f1365])).
fof(f1240, plain, (e21 = op2(h9(e10), e23)), inference(backward_demodulation, [], [f333, f368])).
fof(f368, plain, (op2(e22, e23) = h9(e10)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ((op2(op2(e22, e23), e23) = h9(e11)) & (op2(e22, e23) = h9(e10)) & (e23 = h9(e13)) & (e22 = h9(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax22)).
fof(f333, plain, (e21 = op2(op2(e22, e23), e23)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((e21 = op2(op2(e22, e23), e23)) & (e20 = op2(e22, e23))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax13)).
fof(f3656, plain, (~ spl36_74 | ~ spl36_106), inference(avatar_contradiction_clause, [], [f3655])).
fof(f3655, plain, ($false | (~ spl36_74 | ~ spl36_106)), inference(subsumption_resolution, [], [f3654, f3084])).
fof(f3084, plain, (~ (e21 = h11(e10)) | ~ spl36_106), inference(backward_demodulation, [], [f1246, f1055])).
fof(f1246, plain, ~ (op2(e21, e21) = h11(e10)), inference(backward_demodulation, [], [f256, f376])).
fof(f376, plain, (op2(e23, e21) = h11(e10)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ((h11(e11) = op2(op2(e23, e21), e21)) & (op2(e23, e21) = h11(e10)) & (e21 = h11(e13)) & (e23 = h11(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax24)).
fof(f256, plain, ~ (op2(e21, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f3654, plain, ((e21 = h11(e10)) | ~ spl36_74), inference(forward_demodulation, [], [f376, f919])).
fof(f919, plain, ((e21 = op2(e23, e21)) | ~ spl36_74), inference(avatar_component_clause, [], [f917])).
fof(f917, plain, (spl36_74 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_74])])).
fof(f3637, plain, (~ spl36_41 | ~ spl36_42), inference(avatar_contradiction_clause, [], [f3636])).
fof(f3636, plain, ($false | (~ spl36_41 | ~ spl36_42)), inference(subsumption_resolution, [], [f3635, f294])).
fof(f294, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax7)).
fof(f3635, plain, ((e10 = e11) | (~ spl36_41 | ~ spl36_42)), inference(backward_demodulation, [], [f751, f747])).
fof(f747, plain, ((e10 = op1(e11, e11)) | ~ spl36_41), inference(avatar_component_clause, [], [f745])).
fof(f745, plain, (spl36_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_41])])).
fof(f3630, plain, (~ spl36_67 | ~ spl36_68), inference(avatar_contradiction_clause, [], [f3629])).
fof(f3629, plain, ($false | (~ spl36_67 | ~ spl36_68)), inference(subsumption_resolution, [], [f3628, f305])).
fof(f305, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f3628, plain, ((e22 = e23) | (~ spl36_67 | ~ spl36_68)), inference(backward_demodulation, [], [f893, f889])).
fof(f889, plain, ((e22 = op2(e23, e23)) | ~ spl36_67), inference(avatar_component_clause, [], [f887])).
fof(f887, plain, (spl36_67 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl36_67])])).
fof(f3614, plain, (~ spl36_87 | ~ spl36_88), inference(avatar_contradiction_clause, [], [f3613])).
fof(f3613, plain, ($false | (~ spl36_87 | ~ spl36_88)), inference(subsumption_resolution, [], [f3612, f305])).
fof(f3612, plain, ((e22 = e23) | (~ spl36_87 | ~ spl36_88)), inference(forward_demodulation, [], [f978, f974])).
fof(f978, plain, ((e23 = op2(e22, e22)) | ~ spl36_88), inference(avatar_component_clause, [], [f976])).
fof(f976, plain, (spl36_88 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl36_88])])).
fof(f3600, plain, (~ spl36_87 | ~ spl36_95), inference(avatar_contradiction_clause, [], [f3599])).
fof(f3599, plain, ($false | (~ spl36_87 | ~ spl36_95)), inference(subsumption_resolution, [], [f3597, f3116])).
fof(f3116, plain, (~ (e22 = h7(e10)) | ~ spl36_87), inference(backward_demodulation, [], [f1229, f974])).
fof(f1229, plain, ~ (op2(e22, e22) = h7(e10)), inference(backward_demodulation, [], [f283, f360])).
fof(f360, plain, (op2(e22, e20) = h7(e10)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ((h7(e11) = op2(op2(e22, e20), e20)) & (op2(e22, e20) = h7(e10)) & (e20 = h7(e13)) & (e22 = h7(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax20)).
fof(f283, plain, ~ (op2(e22, e20) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f3597, plain, ((e22 = h7(e10)) | ~ spl36_95), inference(backward_demodulation, [], [f360, f1008])).
fof(f1008, plain, ((e22 = op2(e22, e20)) | ~ spl36_95), inference(avatar_component_clause, [], [f1006])).
fof(f1006, plain, (spl36_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_95])])).
fof(f3584, plain, (spl36_339 | ~ spl36_112), inference(avatar_split_clause, [], [f3583, f1078, f2684])).
fof(f1078, plain, (spl36_112 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_112])])).
fof(f3583, plain, ((e23 = h4(e10)) | ~ spl36_112), inference(backward_demodulation, [], [f348, f1080])).
fof(f1080, plain, ((e23 = op2(e21, e20)) | ~ spl36_112), inference(avatar_component_clause, [], [f1078])).
fof(f3570, plain, (~ spl36_87 | ~ spl36_119), inference(avatar_contradiction_clause, [], [f3569])).
fof(f3569, plain, ($false | (~ spl36_87 | ~ spl36_119)), inference(subsumption_resolution, [], [f3567, f3118])).
fof(f3118, plain, (~ (e22 = h2(e10)) | ~ spl36_87), inference(backward_demodulation, [], [f1200, f974])).
fof(f1200, plain, ~ (op2(e22, e22) = h2(e10)), inference(backward_demodulation, [], [f259, f340])).
fof(f259, plain, ~ (op2(e20, e22) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f3567, plain, ((e22 = h2(e10)) | ~ spl36_119), inference(backward_demodulation, [], [f340, f1110])).
fof(f1110, plain, ((e22 = op2(e20, e22)) | ~ spl36_119), inference(avatar_component_clause, [], [f1108])).
fof(f1108, plain, (spl36_119 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl36_119])])).
fof(f3565, plain, (~ spl36_124 | spl36_393), inference(avatar_contradiction_clause, [], [f3564])).
fof(f3564, plain, ($false | (~ spl36_124 | spl36_393)), inference(subsumption_resolution, [], [f3563, f3054])).
fof(f3054, plain, (~ (e23 = h1(e10)) | spl36_393), inference(avatar_component_clause, [], [f3052])).
fof(f3052, plain, (spl36_393 <=> (e23 = h1(e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_393])])).
fof(f3563, plain, ((e23 = h1(e10)) | ~ spl36_124), inference(backward_demodulation, [], [f336, f1131])).
fof(f1131, plain, ((e23 = op2(e20, e21)) | ~ spl36_124), inference(avatar_component_clause, [], [f1129])).
fof(f1129, plain, (spl36_124 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_124])])).
fof(f3549, plain, (spl36_160 | ~ spl36_97), inference(avatar_split_clause, [], [f3548, f1015, f1437])).
fof(f1437, plain, (spl36_160 <=> (e20 = h6(e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_160])])).
fof(f1015, plain, (spl36_97 <=> (e20 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl36_97])])).
fof(f3548, plain, ((e20 = h6(e10)) | ~ spl36_97), inference(forward_demodulation, [], [f356, f1017])).
fof(f1017, plain, ((e20 = op2(e21, e23)) | ~ spl36_97), inference(avatar_component_clause, [], [f1015])).
fof(f356, plain, (op2(e21, e23) = h6(e10)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ((h6(e11) = op2(op2(e21, e23), e23)) & (op2(e21, e23) = h6(e10)) & (e23 = h6(e13)) & (e21 = h6(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax19)).
fof(f3528, plain, (~ spl36_60 | ~ spl36_56), inference(avatar_split_clause, [], [f3527, f808, f825])).
fof(f825, plain, (spl36_60 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_60])])).
fof(f3527, plain, (~ (e13 = op1(e10, e11)) | ~ spl36_56), inference(forward_demodulation, [], [f225, f810])).
fof(f225, plain, ~ (op1(e10, e11) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax5)).
fof(f3499, plain, (~ spl36_22 | ~ spl36_23), inference(avatar_contradiction_clause, [], [f3498])).
fof(f3498, plain, ($false | (~ spl36_22 | ~ spl36_23)), inference(subsumption_resolution, [], [f3497, f297])).
fof(f297, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f7])).
fof(f3497, plain, ((e11 = e12) | (~ spl36_22 | ~ spl36_23)), inference(backward_demodulation, [], [f670, f666])).
fof(f666, plain, ((e11 = op1(e12, e12)) | ~ spl36_22), inference(avatar_component_clause, [], [f664])).
fof(f664, plain, (spl36_22 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl36_22])])).
fof(f3477, plain, (~ spl36_66 | ~ spl36_68), inference(avatar_contradiction_clause, [], [f3476])).
fof(f3476, plain, ($false | (~ spl36_66 | ~ spl36_68)), inference(subsumption_resolution, [], [f3475, f304])).
fof(f304, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f8])).
fof(f3475, plain, ((e21 = e23) | (~ spl36_66 | ~ spl36_68)), inference(backward_demodulation, [], [f893, f885])).
fof(f885, plain, ((e21 = op2(e23, e23)) | ~ spl36_66), inference(avatar_component_clause, [], [f883])).
fof(f883, plain, (spl36_66 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl36_66])])).
fof(f3470, plain, (~ spl36_73 | ~ spl36_75), inference(avatar_contradiction_clause, [], [f3469])).
fof(f3469, plain, ($false | (~ spl36_73 | ~ spl36_75)), inference(subsumption_resolution, [], [f3468, f301])).
fof(f3468, plain, ((e20 = e22) | (~ spl36_73 | ~ spl36_75)), inference(backward_demodulation, [], [f923, f915])).
fof(f923, plain, ((e22 = op2(e23, e21)) | ~ spl36_75), inference(avatar_component_clause, [], [f921])).
fof(f921, plain, (spl36_75 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_75])])).
fof(f3454, plain, (~ spl36_339 | ~ spl36_96), inference(avatar_split_clause, [], [f3452, f1010, f2684])).
fof(f1010, plain, (spl36_96 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_96])])).
fof(f3452, plain, (~ (e23 = h4(e10)) | ~ spl36_96), inference(backward_demodulation, [], [f1211, f1012])).
fof(f1012, plain, ((e23 = op2(e22, e20)) | ~ spl36_96), inference(avatar_component_clause, [], [f1010])).
fof(f1211, plain, ~ (op2(e22, e20) = h4(e10)), inference(backward_demodulation, [], [f249, f348])).
fof(f249, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f3440, plain, (~ spl36_106 | ~ spl36_110), inference(avatar_contradiction_clause, [], [f3439])).
fof(f3439, plain, ($false | (~ spl36_106 | ~ spl36_110)), inference(subsumption_resolution, [], [f3438, f3088])).
fof(f3088, plain, (~ (e21 = h4(e10)) | ~ spl36_106), inference(backward_demodulation, [], [f1213, f1055])).
fof(f1213, plain, ~ (op2(e21, e21) = h4(e10)), inference(backward_demodulation, [], [f276, f348])).
fof(f276, plain, ~ (op2(e21, e20) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f3438, plain, ((e21 = h4(e10)) | ~ spl36_110), inference(backward_demodulation, [], [f348, f1072])).
fof(f1072, plain, ((e21 = op2(e21, e20)) | ~ spl36_110), inference(avatar_component_clause, [], [f1070])).
fof(f1070, plain, (spl36_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_110])])).
fof(f3436, plain, (~ spl36_393 | ~ spl36_120), inference(avatar_split_clause, [], [f3434, f1112, f3052])).
fof(f3434, plain, (~ (e23 = h1(e10)) | ~ spl36_120), inference(backward_demodulation, [], [f1196, f1114])).
fof(f1196, plain, ~ (op2(e20, e22) = h1(e10)), inference(backward_demodulation, [], [f273, f336])).
fof(f273, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f3432, plain, (~ spl36_106 | ~ spl36_122), inference(avatar_contradiction_clause, [], [f3431])).
fof(f3431, plain, ($false | (~ spl36_106 | ~ spl36_122)), inference(subsumption_resolution, [], [f3430, f3089])).
fof(f3089, plain, (~ (e21 = h1(e10)) | ~ spl36_106), inference(backward_demodulation, [], [f1192, f1055])).
fof(f1192, plain, ~ (op2(e21, e21) = h1(e10)), inference(backward_demodulation, [], [f252, f336])).
fof(f252, plain, ~ (op2(e20, e21) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f3430, plain, ((e21 = h1(e10)) | ~ spl36_122), inference(backward_demodulation, [], [f336, f1123])).
fof(f1123, plain, ((e21 = op2(e20, e21)) | ~ spl36_122), inference(avatar_component_clause, [], [f1121])).
fof(f1121, plain, (spl36_122 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_122])])).
fof(f3417, plain, (spl36_148 | ~ spl36_81), inference(avatar_split_clause, [], [f3416, f947, f1365])).
fof(f3416, plain, ((e20 = h9(e10)) | ~ spl36_81), inference(forward_demodulation, [], [f368, f949])).
fof(f3414, plain, (~ spl36_160 | ~ spl36_81), inference(avatar_split_clause, [], [f3413, f947, f1437])).
fof(f3413, plain, (~ (e20 = h6(e10)) | ~ spl36_81), inference(forward_demodulation, [], [f1222, f949])).
fof(f1222, plain, ~ (op2(e22, e23) = h6(e10)), inference(backward_demodulation, [], [f267, f356])).
fof(f267, plain, ~ (op2(e21, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f3411, plain, (~ spl36_151 | ~ spl36_81), inference(avatar_split_clause, [], [f3410, f947, f1384])).
fof(f1384, plain, (spl36_151 <=> (e20 = h8(e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_151])])).
fof(f3410, plain, (~ (e20 = h8(e10)) | ~ spl36_81), inference(forward_demodulation, [], [f1235, f949])).
fof(f1235, plain, ~ (op2(e22, e23) = h8(e10)), inference(backward_demodulation, [], [f286, f364])).
fof(f364, plain, (op2(e22, e21) = h8(e10)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ((h8(e11) = op2(op2(e22, e21), e21)) & (op2(e22, e21) = h8(e10)) & (e21 = h8(e13)) & (e22 = h8(e12))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax21)).
fof(f286, plain, ~ (op2(e22, e21) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f3407, plain, (~ spl36_166 | ~ spl36_79), inference(avatar_split_clause, [], [f3406, f938, f1472])).
fof(f3406, plain, (~ (e22 = h4(e10)) | ~ spl36_79), inference(forward_demodulation, [], [f1212, f940])).
fof(f1212, plain, ~ (op2(e23, e20) = h4(e10)), inference(backward_demodulation, [], [f250, f348])).
fof(f250, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f3385, plain, (~ spl36_58 | ~ spl36_50), inference(avatar_split_clause, [], [f3384, f783, f817])).
fof(f817, plain, (spl36_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_58])])).
fof(f3384, plain, (~ (e11 = op1(e10, e11)) | ~ spl36_50), inference(forward_demodulation, [], [f226, f785])).
fof(f785, plain, ((e11 = op1(e10, e13)) | ~ spl36_50), inference(avatar_component_clause, [], [f783])).
fof(f226, plain, ~ (op1(e10, e11) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f3381, plain, (~ spl36_55 | ~ spl36_23), inference(avatar_split_clause, [], [f3380, f668, f804])).
fof(f804, plain, (spl36_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl36_55])])).
fof(f3380, plain, (~ (e12 = op1(e10, e12)) | ~ spl36_23), inference(forward_demodulation, [], [f211, f670])).
fof(f211, plain, ~ (op1(e10, e12) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f3377, plain, (~ spl36_54 | ~ spl36_50), inference(avatar_split_clause, [], [f3376, f783, f800])).
fof(f800, plain, (spl36_54 <=> (e11 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl36_54])])).
fof(f3376, plain, (~ (e11 = op1(e10, e12)) | ~ spl36_50), inference(forward_demodulation, [], [f227, f785])).
fof(f227, plain, ~ (op1(e10, e12) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f3369, plain, (~ spl36_50 | spl36_391), inference(avatar_contradiction_clause, [], [f3368])).
fof(f3368, plain, ($false | (~ spl36_50 | spl36_391)), inference(trivial_inequality_removal, [], [f3367])).
fof(f3367, plain, (~ (h1(e11) = h1(e11)) | (~ spl36_50 | spl36_391)), inference(forward_demodulation, [], [f3026, f785])).
fof(f3026, plain, (~ (h1(e11) = h1(op1(e10, e13))) | spl36_391), inference(avatar_component_clause, [], [f3024])).
fof(f3024, plain, (spl36_391 <=> (h1(e11) = h1(op1(e10, e13)))), introduced(avatar_definition, [new_symbols(naming, [spl36_391])])).
fof(f3364, plain, (~ spl36_46 | ~ spl36_42), inference(avatar_split_clause, [], [f3363, f749, f766])).
fof(f766, plain, (spl36_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_46])])).
fof(f3363, plain, (~ (e11 = op1(e11, e10)) | ~ spl36_42), inference(forward_demodulation, [], [f228, f751])).
fof(f228, plain, ~ (op1(e11, e10) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f3360, plain, (~ spl36_47 | ~ spl36_35), inference(avatar_split_clause, [], [f3359, f719, f770])).
fof(f770, plain, (spl36_47 <=> (e12 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_47])])).
fof(f3359, plain, (~ (e12 = op1(e11, e10)) | ~ spl36_35), inference(forward_demodulation, [], [f230, f721])).
fof(f230, plain, ~ (op1(e11, e10) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f3358, plain, (~ spl36_33 | ~ spl36_17), inference(avatar_split_clause, [], [f3357, f643, f711])).
fof(f711, plain, (spl36_33 <=> (e10 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl36_33])])).
fof(f3357, plain, (~ (e10 = op1(e11, e13)) | ~ spl36_17), inference(forward_demodulation, [], [f219, f645])).
fof(f219, plain, ~ (op1(e11, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f3346, plain, (~ spl36_27 | ~ spl36_23), inference(avatar_split_clause, [], [f3345, f668, f685])).
fof(f685, plain, (spl36_27 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_27])])).
fof(f3345, plain, (~ (e12 = op1(e12, e11)) | ~ spl36_23), inference(forward_demodulation, [], [f237, f670])).
fof(f237, plain, ~ (op1(e12, e11) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f3344, plain, (~ spl36_25 | ~ spl36_17), inference(avatar_split_clause, [], [f3343, f643, f677])).
fof(f677, plain, (spl36_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_25])])).
fof(f3343, plain, (~ (e10 = op1(e12, e11)) | ~ spl36_17), inference(forward_demodulation, [], [f238, f645])).
fof(f238, plain, ~ (op1(e12, e11) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f3336, plain, (~ spl36_14 | ~ spl36_6), inference(avatar_split_clause, [], [f3335, f596, f630])).
fof(f630, plain, (spl36_14 <=> (e11 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_14])])).
fof(f3335, plain, (~ (e11 = op1(e13, e10)) | ~ spl36_6), inference(forward_demodulation, [], [f241, f598])).
fof(f241, plain, ~ (op1(e13, e10) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f3305, plain, (~ spl36_4 | ~ spl36_16), inference(avatar_split_clause, [], [f3302, f638, f587])).
fof(f638, plain, (spl36_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_16])])).
fof(f3302, plain, (~ (e13 = op1(e13, e13)) | ~ spl36_16), inference(backward_demodulation, [], [f242, f640])).
fof(f640, plain, ((e13 = op1(e13, e10)) | ~ spl36_16), inference(avatar_component_clause, [], [f638])).
fof(f242, plain, ~ (op1(e13, e10) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f3299, plain, (~ spl36_17 | ~ spl36_18), inference(avatar_contradiction_clause, [], [f3298])).
fof(f3298, plain, ($false | (~ spl36_17 | ~ spl36_18)), inference(subsumption_resolution, [], [f3297, f294])).
fof(f3297, plain, ((e10 = e11) | (~ spl36_17 | ~ spl36_18)), inference(backward_demodulation, [], [f649, f645])).
fof(f649, plain, ((e11 = op1(e12, e13)) | ~ spl36_18), inference(avatar_component_clause, [], [f647])).
fof(f647, plain, (spl36_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl36_18])])).
fof(f3268, plain, (~ spl36_4 | ~ spl36_36), inference(avatar_split_clause, [], [f3266, f723, f587])).
fof(f723, plain, (spl36_36 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl36_36])])).
fof(f3266, plain, (~ (e13 = op1(e13, e13)) | ~ spl36_36), inference(backward_demodulation, [], [f220, f725])).
fof(f725, plain, ((e13 = op1(e11, e13)) | ~ spl36_36), inference(avatar_component_clause, [], [f723])).
fof(f220, plain, ~ (op1(e11, e13) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f3258, plain, (~ spl36_34 | ~ spl36_42), inference(avatar_split_clause, [], [f3254, f749, f715])).
fof(f715, plain, (spl36_34 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl36_34])])).
fof(f3254, plain, (~ (e11 = op1(e11, e13)) | ~ spl36_42), inference(backward_demodulation, [], [f232, f751])).
fof(f232, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f3257, plain, (~ spl36_38 | ~ spl36_42), inference(avatar_split_clause, [], [f3253, f749, f732])).
fof(f732, plain, (spl36_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl36_38])])).
fof(f3253, plain, (~ (e11 = op1(e11, e12)) | ~ spl36_42), inference(backward_demodulation, [], [f231, f751])).
fof(f231, plain, ~ (op1(e11, e11) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f3255, plain, (~ spl36_26 | ~ spl36_42), inference(avatar_split_clause, [], [f3251, f749, f681])).
fof(f681, plain, (spl36_26 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_26])])).
fof(f3251, plain, (~ (e11 = op1(e12, e11)) | ~ spl36_42), inference(backward_demodulation, [], [f207, f751])).
fof(f207, plain, ~ (op1(e11, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f3206, plain, (~ spl36_53 | ~ spl36_61), inference(avatar_split_clause, [], [f3190, f830, f796])).
fof(f796, plain, (spl36_53 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl36_53])])).
fof(f3190, plain, (~ (e10 = op1(e10, e12)) | ~ spl36_61), inference(backward_demodulation, [], [f223, f832])).
fof(f223, plain, ~ (op1(e10, e10) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f3205, plain, (~ spl36_57 | ~ spl36_61), inference(avatar_split_clause, [], [f3189, f830, f813])).
fof(f813, plain, (spl36_57 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl36_57])])).
fof(f3189, plain, (~ (e10 = op1(e10, e11)) | ~ spl36_61), inference(backward_demodulation, [], [f222, f832])).
fof(f222, plain, ~ (op1(e10, e10) = op1(e10, e11)), inference(cnf_transformation, [], [f5])).
fof(f3204, plain, (~ spl36_13 | ~ spl36_61), inference(avatar_split_clause, [], [f3188, f830, f626])).
fof(f626, plain, (spl36_13 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_13])])).
fof(f3188, plain, (~ (e10 = op1(e13, e10)) | ~ spl36_61), inference(backward_demodulation, [], [f200, f832])).
fof(f200, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f3202, plain, (~ spl36_45 | ~ spl36_61), inference(avatar_split_clause, [], [f3186, f830, f762])).
fof(f762, plain, (spl36_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_45])])).
fof(f3186, plain, (~ (e10 = op1(e11, e10)) | ~ spl36_61), inference(backward_demodulation, [], [f198, f832])).
fof(f198, plain, ~ (op1(e10, e10) = op1(e11, e10)), inference(cnf_transformation, [], [f5])).
fof(f3180, plain, (~ spl36_249 | ~ spl36_68), inference(avatar_split_clause, [], [f3174, f891, f2080])).
fof(f2080, plain, (spl36_249 <=> (e23 = h9(e10))), introduced(avatar_definition, [new_symbols(naming, [spl36_249])])).
fof(f3174, plain, (~ (e23 = h9(e10)) | ~ spl36_68), inference(backward_demodulation, [], [f1238, f893])).
fof(f1238, plain, ~ (op2(e23, e23) = h9(e10)), inference(backward_demodulation, [], [f269, f368])).
fof(f269, plain, ~ (op2(e22, e23) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f3146, plain, (~ spl36_78 | ~ spl36_79), inference(avatar_contradiction_clause, [], [f3145])).
fof(f3145, plain, ($false | (~ spl36_78 | ~ spl36_79)), inference(subsumption_resolution, [], [f3144, f303])).
fof(f303, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f8])).
fof(f3144, plain, ((e21 = e22) | (~ spl36_78 | ~ spl36_79)), inference(backward_demodulation, [], [f940, f936])).
fof(f936, plain, ((e21 = op2(e23, e20)) | ~ spl36_78), inference(avatar_component_clause, [], [f934])).
fof(f934, plain, (spl36_78 <=> (e21 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_78])])).
fof(f3124, plain, (spl36_249 | ~ spl36_84), inference(avatar_split_clause, [], [f3119, f959, f2080])).
fof(f959, plain, (spl36_84 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl36_84])])).
fof(f3119, plain, ((e23 = h9(e10)) | ~ spl36_84), inference(backward_demodulation, [], [f368, f961])).
fof(f961, plain, ((e23 = op2(e22, e23)) | ~ spl36_84), inference(avatar_component_clause, [], [f959])).
fof(f3112, plain, (spl36_151 | ~ spl36_89), inference(avatar_split_clause, [], [f3109, f981, f1384])).
fof(f981, plain, (spl36_89 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_89])])).
fof(f3109, plain, ((e20 = h8(e10)) | ~ spl36_89), inference(backward_demodulation, [], [f364, f983])).
fof(f983, plain, ((e20 = op2(e22, e21)) | ~ spl36_89), inference(avatar_component_clause, [], [f981])).
fof(f3108, plain, (~ spl36_93 | ~ spl36_125), inference(avatar_contradiction_clause, [], [f3107])).
fof(f3107, plain, ($false | (~ spl36_93 | ~ spl36_125)), inference(subsumption_resolution, [], [f3105, f3057])).
fof(f3057, plain, (~ (e20 = h7(e10)) | ~ spl36_125), inference(backward_demodulation, [], [f1226, f1136])).
fof(f1226, plain, ~ (op2(e20, e20) = h7(e10)), inference(backward_demodulation, [], [f247, f360])).
fof(f247, plain, ~ (op2(e20, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f3105, plain, ((e20 = h7(e10)) | ~ spl36_93), inference(backward_demodulation, [], [f360, f1000])).
fof(f1000, plain, ((e20 = op2(e22, e20)) | ~ spl36_93), inference(avatar_component_clause, [], [f998])).
fof(f998, plain, (spl36_93 <=> (e20 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_93])])).
fof(f3083, plain, (~ spl36_109 | ~ spl36_125), inference(avatar_contradiction_clause, [], [f3082])).
fof(f3082, plain, ($false | (~ spl36_109 | ~ spl36_125)), inference(subsumption_resolution, [], [f3081, f3058])).
fof(f3058, plain, (~ (e20 = h4(e10)) | ~ spl36_125), inference(backward_demodulation, [], [f1210, f1136])).
fof(f1210, plain, ~ (op2(e20, e20) = h4(e10)), inference(backward_demodulation, [], [f246, f348])).
fof(f246, plain, ~ (op2(e20, e20) = op2(e21, e20)), inference(cnf_transformation, [], [f6])).
fof(f3081, plain, ((e20 = h4(e10)) | ~ spl36_109), inference(backward_demodulation, [], [f348, f1068])).
fof(f1068, plain, ((e20 = op2(e21, e20)) | ~ spl36_109), inference(avatar_component_clause, [], [f1066])).
fof(f1066, plain, (spl36_109 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl36_109])])).
fof(f3080, plain, (~ spl36_114 | ~ spl36_115), inference(avatar_contradiction_clause, [], [f3079])).
fof(f3079, plain, ($false | (~ spl36_114 | ~ spl36_115)), inference(subsumption_resolution, [], [f3078, f303])).
fof(f3078, plain, ((e21 = e22) | (~ spl36_114 | ~ spl36_115)), inference(backward_demodulation, [], [f1093, f1089])).
fof(f1093, plain, ((e22 = op2(e20, e23)) | ~ spl36_115), inference(avatar_component_clause, [], [f1091])).
fof(f1091, plain, (spl36_115 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl36_115])])).
fof(f3068, plain, (~ spl36_117 | ~ spl36_125), inference(avatar_contradiction_clause, [], [f3067])).
fof(f3067, plain, ($false | (~ spl36_117 | ~ spl36_125)), inference(subsumption_resolution, [], [f3065, f3060])).
fof(f3060, plain, (~ (e20 = h2(e10)) | ~ spl36_125), inference(backward_demodulation, [], [f1202, f1136])).
fof(f1202, plain, ~ (op2(e20, e20) = h2(e10)), inference(backward_demodulation, [], [f271, f340])).
fof(f271, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f3065, plain, ((e20 = h2(e10)) | ~ spl36_117), inference(backward_demodulation, [], [f340, f1102])).
fof(f1102, plain, ((e20 = op2(e20, e22)) | ~ spl36_117), inference(avatar_component_clause, [], [f1100])).
fof(f1100, plain, (spl36_117 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl36_117])])).
fof(f3064, plain, (~ spl36_121 | ~ spl36_125), inference(avatar_contradiction_clause, [], [f3063])).
fof(f3063, plain, ($false | (~ spl36_121 | ~ spl36_125)), inference(subsumption_resolution, [], [f3062, f3061])).
fof(f3061, plain, (~ (e20 = h1(e10)) | ~ spl36_125), inference(backward_demodulation, [], [f1195, f1136])).
fof(f1195, plain, ~ (op2(e20, e20) = h1(e10)), inference(backward_demodulation, [], [f270, f336])).
fof(f270, plain, ~ (op2(e20, e20) = op2(e20, e21)), inference(cnf_transformation, [], [f6])).
fof(f3062, plain, ((e20 = h1(e10)) | ~ spl36_121), inference(backward_demodulation, [], [f336, f1119])).
fof(f1119, plain, ((e20 = op2(e20, e21)) | ~ spl36_121), inference(avatar_component_clause, [], [f1117])).
fof(f1117, plain, (spl36_121 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl36_121])])).
fof(f3031, plain, (~ spl36_376 | ~ spl36_377 | ~ spl36_378 | ~ spl36_379 | spl36_176 | ~ spl36_380 | ~ spl36_381 | ~ spl36_382 | ~ spl36_383 | ~ spl36_384 | ~ spl36_385 | ~ spl36_386 | ~ spl36_387 | ~ spl36_388 | ~ spl36_389 | ~ spl36_390 | ~ spl36_391 | ~ spl36_392), inference(avatar_split_clause, [], [f2962, f3028, f3024, f3020, f3016, f3012, f3008, f3004, f3000, f2996, f2992, f2988, f2984, f2980, f1535, f2976, f2972, f2968, f2964])).
fof(f1535, plain, (spl36_176 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl36_176])])).
fof(f2962, plain, (~ (h1(op1(e10, e12)) = op2(h1(e10), e20)) | ~ (h1(e11) = h1(op1(e10, e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e13)) = op2(h1(e11), e21)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(op1(e12, e11)) = op2(e20, h1(e11))) | ~ (op2(e20, e20) = h1(op1(e12, e12))) | ~ (h1(e10) = h1(op1(e12, e13))) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2961, f334])).
fof(f2961, plain, (~ (h1(e11) = h1(op1(e10, e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e13)) = op2(h1(e11), e21)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(op1(e12, e11)) = op2(e20, h1(e11))) | ~ (op2(e20, e20) = h1(op1(e12, e12))) | ~ (h1(e10) = h1(op1(e12, e13))) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2960, f1198])).
fof(f2960, plain, (~ (h1(op1(e10, e13)) = op2(h1(e10), e21)) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e13)) = op2(h1(e11), e21)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(op1(e12, e11)) = op2(e20, h1(e11))) | ~ (op2(e20, e20) = h1(op1(e12, e12))) | ~ (h1(e10) = h1(op1(e12, e13))) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2959, f335])).
fof(f2959, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e13)) = op2(h1(e11), e21)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(op1(e12, e11)) = op2(e20, h1(e11))) | ~ (op2(e20, e20) = h1(op1(e12, e12))) | ~ (h1(e10) = h1(op1(e12, e13))) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2958, f334])).
fof(f2958, plain, (~ (h1(op1(e11, e13)) = op2(h1(e11), e21)) | ~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(op1(e12, e11)) = op2(e20, h1(e11))) | ~ (op2(e20, e20) = h1(op1(e12, e12))) | ~ (h1(e10) = h1(op1(e12, e13))) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2957, f335])).
fof(f2957, plain, (~ (h1(op1(e12, e10)) = op2(e20, h1(e10))) | ~ (h1(op1(e12, e11)) = op2(e20, h1(e11))) | ~ (op2(e20, e20) = h1(op1(e12, e12))) | ~ (h1(e10) = h1(op1(e12, e13))) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2956, f334])).
fof(f2956, plain, (~ (h1(op1(e12, e11)) = op2(e20, h1(e11))) | ~ (op2(e20, e20) = h1(op1(e12, e12))) | ~ (h1(e10) = h1(op1(e12, e13))) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2955, f334])).
fof(f2955, plain, (~ (op2(e20, e20) = h1(op1(e12, e12))) | ~ (h1(e10) = h1(op1(e12, e13))) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2954, f334])).
fof(f2954, plain, (~ (h1(e10) = h1(op1(e12, e13))) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2953, f336])).
fof(f2953, plain, (~ (op2(e20, e21) = h1(op1(e12, e13))) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2952, f334])).
fof(f2952, plain, (~ (h1(op1(e12, e13)) = op2(h1(e12), e21)) | ~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2951, f335])).
fof(f2951, plain, (~ (h1(op1(e13, e10)) = op2(e21, h1(e10))) | ~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2950, f335])).
fof(f2950, plain, (~ (h1(op1(e13, e11)) = op2(e21, h1(e11))) | ~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2949, f335])).
fof(f2949, plain, (~ (h4(e10) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2948, f348])).
fof(f2948, plain, (~ (op2(e21, e20) = h1(op1(e13, e12))) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2947, f335])).
fof(f2947, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), e20)) | ~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2946, f334])).
fof(f2946, plain, (~ (op2(e21, e21) = h1(op1(e13, e13))) | ~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f2945, f335])).
fof(f2945, plain, (~ (e23 = h1(e11)) | sP2 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(subsumption_resolution, [], [f2944, f1550])).
fof(f1550, plain, ~ sP0, inference(subsumption_resolution, [], [f524, f334])).
fof(f524, plain, (~ (e20 = h1(e12)) | ~ sP0), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP0), inference(usedef, [], [e29])).
fof(e29, plain, (sP0 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f2944, plain, (~ (e23 = h1(e11)) | sP2 | sP0 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(subsumption_resolution, [], [f527, f1548])).
fof(f1548, plain, ~ sP1, inference(subsumption_resolution, [], [f521, f335])).
fof(f521, plain, (~ (e21 = h1(e13)) | ~ sP1), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP1), inference(usedef, [], [e30])).
fof(e30, plain, (sP1 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f527, plain, (~ (e23 = h1(e11)) | sP2 | sP1 | sP0 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(cnf_transformation, [], [f65])).
fof(f65, plain, (((~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | sP35 | sP34 | sP33 | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | sP32 | sP31 | sP30 | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | sP29 | sP28 | sP27 | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | sP26 | sP25 | sP24 | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | sP23 | sP22 | sP21 | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | sP20 | sP19 | sP18 | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | sP17 | sP16 | sP15 | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | sP14 | sP13 | sP12 | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP11 | sP10 | sP9 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP8 | sP7 | sP6 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP5 | sP4 | sP3 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP2 | sP1 | sP0 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f28, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29])).
fof(f31, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(usedef, [], [e31])).
fof(e31, plain, (sP2 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f32, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP3), inference(usedef, [], [e32])).
fof(e32, plain, (sP3 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f33, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP4), inference(usedef, [], [e33])).
fof(e33, plain, (sP4 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f34, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP5), inference(usedef, [], [e34])).
fof(e34, plain, (sP5 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f35, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP6), inference(usedef, [], [e35])).
fof(e35, plain, (sP6 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f36, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP7), inference(usedef, [], [e36])).
fof(e36, plain, (sP7 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f37, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP8), inference(usedef, [], [e37])).
fof(e37, plain, (sP8 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f38, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP9), inference(usedef, [], [e38])).
fof(e38, plain, (sP9 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f39, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP10), inference(usedef, [], [e39])).
fof(e39, plain, (sP10 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f40, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP11), inference(usedef, [], [e40])).
fof(e40, plain, (sP11 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f41, plain, ((~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ sP12), inference(usedef, [], [e41])).
fof(e41, plain, (sP12 <=> (~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f42, plain, ((~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | ~ sP13), inference(usedef, [], [e42])).
fof(e42, plain, (sP13 <=> (~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f43, plain, ((~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | ~ sP14), inference(usedef, [], [e43])).
fof(e43, plain, (sP14 <=> (~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f44, plain, ((~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ sP15), inference(usedef, [], [e44])).
fof(e44, plain, (sP15 <=> (~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f45, plain, ((~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | ~ sP16), inference(usedef, [], [e45])).
fof(e45, plain, (sP16 <=> (~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f46, plain, ((~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | ~ sP17), inference(usedef, [], [e46])).
fof(e46, plain, (sP17 <=> (~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f47, plain, ((~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ sP18), inference(usedef, [], [e47])).
fof(e47, plain, (sP18 <=> (~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f48, plain, ((~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | ~ sP19), inference(usedef, [], [e48])).
fof(e48, plain, (sP19 <=> (~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f49, plain, ((~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | ~ sP20), inference(usedef, [], [e49])).
fof(e49, plain, (sP20 <=> (~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f50, plain, ((~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ sP21), inference(usedef, [], [e50])).
fof(e50, plain, (sP21 <=> (~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f51, plain, ((~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | ~ sP22), inference(usedef, [], [e51])).
fof(e51, plain, (sP22 <=> (~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f52, plain, ((~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | ~ sP23), inference(usedef, [], [e52])).
fof(e52, plain, (sP23 <=> (~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f53, plain, ((~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ sP24), inference(usedef, [], [e53])).
fof(e53, plain, (sP24 <=> (~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f54, plain, ((~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | ~ sP25), inference(usedef, [], [e54])).
fof(e54, plain, (sP25 <=> (~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f55, plain, ((~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | ~ sP26), inference(usedef, [], [e55])).
fof(e55, plain, (sP26 <=> (~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f56, plain, ((~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ sP27), inference(usedef, [], [e56])).
fof(e56, plain, (sP27 <=> (~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f57, plain, ((~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | ~ sP28), inference(usedef, [], [e57])).
fof(e57, plain, (sP28 <=> (~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f58, plain, ((~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | ~ sP29), inference(usedef, [], [e58])).
fof(e58, plain, (sP29 <=> (~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f59, plain, ((~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ sP30), inference(usedef, [], [e59])).
fof(e59, plain, (sP30 <=> (~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f60, plain, ((~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | ~ sP31), inference(usedef, [], [e60])).
fof(e60, plain, (sP31 <=> (~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f61, plain, ((~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | ~ sP32), inference(usedef, [], [e61])).
fof(e61, plain, (sP32 <=> (~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f62, plain, ((~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ sP33), inference(usedef, [], [e62])).
fof(e62, plain, (sP33 <=> (~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f63, plain, ((~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | ~ sP34), inference(usedef, [], [e63])).
fof(e63, plain, (sP34 <=> (~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f64, plain, ((~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | ~ sP35), inference(usedef, [], [e64])).
fof(e64, plain, (sP35 <=> (~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f28, plain, (((~ (e23 = h12(e13)) & ~ (e23 = h12(e12)) & ~ (e23 = h12(e11)) & ~ (e23 = h12(e10))) | (~ (e22 = h12(e13)) & ~ (e22 = h12(e12)) & ~ (e22 = h12(e11)) & ~ (e22 = h12(e10))) | (~ (e21 = h12(e13)) & ~ (e21 = h12(e12)) & ~ (e21 = h12(e11)) & ~ (e21 = h12(e10))) | (~ (e20 = h12(e13)) & ~ (e20 = h12(e12)) & ~ (e20 = h12(e11)) & ~ (e20 = h12(e10))) | ~ (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) | ~ (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) | ~ (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) | ~ (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) | ~ (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) | ~ (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) | ~ (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) | ~ (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) | ~ (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) | ~ (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) | ~ (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) | ~ (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) | ~ (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) | ~ (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) | ~ (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) | ~ (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) & ((~ (e23 = h11(e13)) & ~ (e23 = h11(e12)) & ~ (e23 = h11(e11)) & ~ (e23 = h11(e10))) | (~ (e22 = h11(e13)) & ~ (e22 = h11(e12)) & ~ (e22 = h11(e11)) & ~ (e22 = h11(e10))) | (~ (e21 = h11(e13)) & ~ (e21 = h11(e12)) & ~ (e21 = h11(e11)) & ~ (e21 = h11(e10))) | (~ (e20 = h11(e13)) & ~ (e20 = h11(e12)) & ~ (e20 = h11(e11)) & ~ (e20 = h11(e10))) | ~ (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) | ~ (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) | ~ (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) | ~ (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) | ~ (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) | ~ (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) | ~ (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) | ~ (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) | ~ (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) | ~ (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) | ~ (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) | ~ (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) | ~ (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) | ~ (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) | ~ (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) | ~ (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) & ((~ (e23 = h10(e13)) & ~ (e23 = h10(e12)) & ~ (e23 = h10(e11)) & ~ (e23 = h10(e10))) | (~ (e22 = h10(e13)) & ~ (e22 = h10(e12)) & ~ (e22 = h10(e11)) & ~ (e22 = h10(e10))) | (~ (e21 = h10(e13)) & ~ (e21 = h10(e12)) & ~ (e21 = h10(e11)) & ~ (e21 = h10(e10))) | (~ (e20 = h10(e13)) & ~ (e20 = h10(e12)) & ~ (e20 = h10(e11)) & ~ (e20 = h10(e10))) | ~ (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) | ~ (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) | ~ (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) | ~ (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) | ~ (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) | ~ (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) | ~ (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) | ~ (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) | ~ (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) | ~ (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) | ~ (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) | ~ (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) | ~ (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) | ~ (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) | ~ (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) | ~ (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) & ((~ (e23 = h9(e13)) & ~ (e23 = h9(e12)) & ~ (e23 = h9(e11)) & ~ (e23 = h9(e10))) | (~ (e22 = h9(e13)) & ~ (e22 = h9(e12)) & ~ (e22 = h9(e11)) & ~ (e22 = h9(e10))) | (~ (e21 = h9(e13)) & ~ (e21 = h9(e12)) & ~ (e21 = h9(e11)) & ~ (e21 = h9(e10))) | (~ (e20 = h9(e13)) & ~ (e20 = h9(e12)) & ~ (e20 = h9(e11)) & ~ (e20 = h9(e10))) | ~ (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) | ~ (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) | ~ (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) | ~ (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) | ~ (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) | ~ (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) | ~ (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) | ~ (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) | ~ (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) | ~ (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) | ~ (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) | ~ (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) | ~ (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) | ~ (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) | ~ (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) | ~ (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) & ((~ (e23 = h8(e13)) & ~ (e23 = h8(e12)) & ~ (e23 = h8(e11)) & ~ (e23 = h8(e10))) | (~ (e22 = h8(e13)) & ~ (e22 = h8(e12)) & ~ (e22 = h8(e11)) & ~ (e22 = h8(e10))) | (~ (e21 = h8(e13)) & ~ (e21 = h8(e12)) & ~ (e21 = h8(e11)) & ~ (e21 = h8(e10))) | (~ (e20 = h8(e13)) & ~ (e20 = h8(e12)) & ~ (e20 = h8(e11)) & ~ (e20 = h8(e10))) | ~ (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) | ~ (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) | ~ (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) | ~ (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) | ~ (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) | ~ (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) | ~ (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) | ~ (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) | ~ (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) | ~ (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) | ~ (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) | ~ (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) | ~ (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) | ~ (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) | ~ (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) | ~ (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) & ((~ (e23 = h7(e13)) & ~ (e23 = h7(e12)) & ~ (e23 = h7(e11)) & ~ (e23 = h7(e10))) | (~ (e22 = h7(e13)) & ~ (e22 = h7(e12)) & ~ (e22 = h7(e11)) & ~ (e22 = h7(e10))) | (~ (e21 = h7(e13)) & ~ (e21 = h7(e12)) & ~ (e21 = h7(e11)) & ~ (e21 = h7(e10))) | (~ (e20 = h7(e13)) & ~ (e20 = h7(e12)) & ~ (e20 = h7(e11)) & ~ (e20 = h7(e10))) | ~ (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) | ~ (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) | ~ (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) | ~ (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) | ~ (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) | ~ (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) | ~ (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) | ~ (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) | ~ (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) | ~ (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) | ~ (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) | ~ (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) | ~ (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) | ~ (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) | ~ (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) | ~ (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) & ((~ (e23 = h6(e13)) & ~ (e23 = h6(e12)) & ~ (e23 = h6(e11)) & ~ (e23 = h6(e10))) | (~ (e22 = h6(e13)) & ~ (e22 = h6(e12)) & ~ (e22 = h6(e11)) & ~ (e22 = h6(e10))) | (~ (e21 = h6(e13)) & ~ (e21 = h6(e12)) & ~ (e21 = h6(e11)) & ~ (e21 = h6(e10))) | (~ (e20 = h6(e13)) & ~ (e20 = h6(e12)) & ~ (e20 = h6(e11)) & ~ (e20 = h6(e10))) | ~ (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) | ~ (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) | ~ (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) | ~ (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) | ~ (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) | ~ (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) | ~ (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) | ~ (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) | ~ (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) | ~ (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) | ~ (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) | ~ (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) | ~ (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) | ~ (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) | ~ (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) | ~ (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) & ((~ (e23 = h5(e13)) & ~ (e23 = h5(e12)) & ~ (e23 = h5(e11)) & ~ (e23 = h5(e10))) | (~ (e22 = h5(e13)) & ~ (e22 = h5(e12)) & ~ (e22 = h5(e11)) & ~ (e22 = h5(e10))) | (~ (e21 = h5(e13)) & ~ (e21 = h5(e12)) & ~ (e21 = h5(e11)) & ~ (e21 = h5(e10))) | (~ (e20 = h5(e13)) & ~ (e20 = h5(e12)) & ~ (e20 = h5(e11)) & ~ (e20 = h5(e10))) | ~ (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) | ~ (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) | ~ (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) | ~ (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) | ~ (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) | ~ (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) | ~ (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) | ~ (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) | ~ (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) | ~ (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) | ~ (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) | ~ (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) | ~ (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) | ~ (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) | ~ (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) | ~ (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) & ((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((((e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f26])).
fof(f26, plain, ~ ((((e23 = h12(e13)) | (e23 = h12(e12)) | (e23 = h12(e11)) | (e23 = h12(e10))) & ((e22 = h12(e13)) | (e22 = h12(e12)) | (e22 = h12(e11)) | (e22 = h12(e10))) & ((e21 = h12(e13)) | (e21 = h12(e12)) | (e21 = h12(e11)) | (e21 = h12(e10))) & ((e20 = h12(e13)) | (e20 = h12(e12)) | (e20 = h12(e11)) | (e20 = h12(e10))) & (h12(op1(e13, e13)) = op2(h12(e13), h12(e13))) & (h12(op1(e13, e12)) = op2(h12(e13), h12(e12))) & (h12(op1(e13, e11)) = op2(h12(e13), h12(e11))) & (h12(op1(e13, e10)) = op2(h12(e13), h12(e10))) & (h12(op1(e12, e13)) = op2(h12(e12), h12(e13))) & (h12(op1(e12, e12)) = op2(h12(e12), h12(e12))) & (h12(op1(e12, e11)) = op2(h12(e12), h12(e11))) & (h12(op1(e12, e10)) = op2(h12(e12), h12(e10))) & (h12(op1(e11, e13)) = op2(h12(e11), h12(e13))) & (h12(op1(e11, e12)) = op2(h12(e11), h12(e12))) & (h12(op1(e11, e11)) = op2(h12(e11), h12(e11))) & (h12(op1(e11, e10)) = op2(h12(e11), h12(e10))) & (h12(op1(e10, e13)) = op2(h12(e10), h12(e13))) & (h12(op1(e10, e12)) = op2(h12(e10), h12(e12))) & (h12(op1(e10, e11)) = op2(h12(e10), h12(e11))) & (h12(op1(e10, e10)) = op2(h12(e10), h12(e10)))) | (((e23 = h11(e13)) | (e23 = h11(e12)) | (e23 = h11(e11)) | (e23 = h11(e10))) & ((e22 = h11(e13)) | (e22 = h11(e12)) | (e22 = h11(e11)) | (e22 = h11(e10))) & ((e21 = h11(e13)) | (e21 = h11(e12)) | (e21 = h11(e11)) | (e21 = h11(e10))) & ((e20 = h11(e13)) | (e20 = h11(e12)) | (e20 = h11(e11)) | (e20 = h11(e10))) & (h11(op1(e13, e13)) = op2(h11(e13), h11(e13))) & (h11(op1(e13, e12)) = op2(h11(e13), h11(e12))) & (h11(op1(e13, e11)) = op2(h11(e13), h11(e11))) & (h11(op1(e13, e10)) = op2(h11(e13), h11(e10))) & (h11(op1(e12, e13)) = op2(h11(e12), h11(e13))) & (h11(op1(e12, e12)) = op2(h11(e12), h11(e12))) & (h11(op1(e12, e11)) = op2(h11(e12), h11(e11))) & (h11(op1(e12, e10)) = op2(h11(e12), h11(e10))) & (h11(op1(e11, e13)) = op2(h11(e11), h11(e13))) & (h11(op1(e11, e12)) = op2(h11(e11), h11(e12))) & (h11(op1(e11, e11)) = op2(h11(e11), h11(e11))) & (h11(op1(e11, e10)) = op2(h11(e11), h11(e10))) & (h11(op1(e10, e13)) = op2(h11(e10), h11(e13))) & (h11(op1(e10, e12)) = op2(h11(e10), h11(e12))) & (h11(op1(e10, e11)) = op2(h11(e10), h11(e11))) & (h11(op1(e10, e10)) = op2(h11(e10), h11(e10)))) | (((e23 = h10(e13)) | (e23 = h10(e12)) | (e23 = h10(e11)) | (e23 = h10(e10))) & ((e22 = h10(e13)) | (e22 = h10(e12)) | (e22 = h10(e11)) | (e22 = h10(e10))) & ((e21 = h10(e13)) | (e21 = h10(e12)) | (e21 = h10(e11)) | (e21 = h10(e10))) & ((e20 = h10(e13)) | (e20 = h10(e12)) | (e20 = h10(e11)) | (e20 = h10(e10))) & (h10(op1(e13, e13)) = op2(h10(e13), h10(e13))) & (h10(op1(e13, e12)) = op2(h10(e13), h10(e12))) & (h10(op1(e13, e11)) = op2(h10(e13), h10(e11))) & (h10(op1(e13, e10)) = op2(h10(e13), h10(e10))) & (h10(op1(e12, e13)) = op2(h10(e12), h10(e13))) & (h10(op1(e12, e12)) = op2(h10(e12), h10(e12))) & (h10(op1(e12, e11)) = op2(h10(e12), h10(e11))) & (h10(op1(e12, e10)) = op2(h10(e12), h10(e10))) & (h10(op1(e11, e13)) = op2(h10(e11), h10(e13))) & (h10(op1(e11, e12)) = op2(h10(e11), h10(e12))) & (h10(op1(e11, e11)) = op2(h10(e11), h10(e11))) & (h10(op1(e11, e10)) = op2(h10(e11), h10(e10))) & (h10(op1(e10, e13)) = op2(h10(e10), h10(e13))) & (h10(op1(e10, e12)) = op2(h10(e10), h10(e12))) & (h10(op1(e10, e11)) = op2(h10(e10), h10(e11))) & (h10(op1(e10, e10)) = op2(h10(e10), h10(e10)))) | (((e23 = h9(e13)) | (e23 = h9(e12)) | (e23 = h9(e11)) | (e23 = h9(e10))) & ((e22 = h9(e13)) | (e22 = h9(e12)) | (e22 = h9(e11)) | (e22 = h9(e10))) & ((e21 = h9(e13)) | (e21 = h9(e12)) | (e21 = h9(e11)) | (e21 = h9(e10))) & ((e20 = h9(e13)) | (e20 = h9(e12)) | (e20 = h9(e11)) | (e20 = h9(e10))) & (h9(op1(e13, e13)) = op2(h9(e13), h9(e13))) & (h9(op1(e13, e12)) = op2(h9(e13), h9(e12))) & (h9(op1(e13, e11)) = op2(h9(e13), h9(e11))) & (h9(op1(e13, e10)) = op2(h9(e13), h9(e10))) & (h9(op1(e12, e13)) = op2(h9(e12), h9(e13))) & (h9(op1(e12, e12)) = op2(h9(e12), h9(e12))) & (h9(op1(e12, e11)) = op2(h9(e12), h9(e11))) & (h9(op1(e12, e10)) = op2(h9(e12), h9(e10))) & (h9(op1(e11, e13)) = op2(h9(e11), h9(e13))) & (h9(op1(e11, e12)) = op2(h9(e11), h9(e12))) & (h9(op1(e11, e11)) = op2(h9(e11), h9(e11))) & (h9(op1(e11, e10)) = op2(h9(e11), h9(e10))) & (h9(op1(e10, e13)) = op2(h9(e10), h9(e13))) & (h9(op1(e10, e12)) = op2(h9(e10), h9(e12))) & (h9(op1(e10, e11)) = op2(h9(e10), h9(e11))) & (h9(op1(e10, e10)) = op2(h9(e10), h9(e10)))) | (((e23 = h8(e13)) | (e23 = h8(e12)) | (e23 = h8(e11)) | (e23 = h8(e10))) & ((e22 = h8(e13)) | (e22 = h8(e12)) | (e22 = h8(e11)) | (e22 = h8(e10))) & ((e21 = h8(e13)) | (e21 = h8(e12)) | (e21 = h8(e11)) | (e21 = h8(e10))) & ((e20 = h8(e13)) | (e20 = h8(e12)) | (e20 = h8(e11)) | (e20 = h8(e10))) & (h8(op1(e13, e13)) = op2(h8(e13), h8(e13))) & (h8(op1(e13, e12)) = op2(h8(e13), h8(e12))) & (h8(op1(e13, e11)) = op2(h8(e13), h8(e11))) & (h8(op1(e13, e10)) = op2(h8(e13), h8(e10))) & (h8(op1(e12, e13)) = op2(h8(e12), h8(e13))) & (h8(op1(e12, e12)) = op2(h8(e12), h8(e12))) & (h8(op1(e12, e11)) = op2(h8(e12), h8(e11))) & (h8(op1(e12, e10)) = op2(h8(e12), h8(e10))) & (h8(op1(e11, e13)) = op2(h8(e11), h8(e13))) & (h8(op1(e11, e12)) = op2(h8(e11), h8(e12))) & (h8(op1(e11, e11)) = op2(h8(e11), h8(e11))) & (h8(op1(e11, e10)) = op2(h8(e11), h8(e10))) & (h8(op1(e10, e13)) = op2(h8(e10), h8(e13))) & (h8(op1(e10, e12)) = op2(h8(e10), h8(e12))) & (h8(op1(e10, e11)) = op2(h8(e10), h8(e11))) & (h8(op1(e10, e10)) = op2(h8(e10), h8(e10)))) | (((e23 = h7(e13)) | (e23 = h7(e12)) | (e23 = h7(e11)) | (e23 = h7(e10))) & ((e22 = h7(e13)) | (e22 = h7(e12)) | (e22 = h7(e11)) | (e22 = h7(e10))) & ((e21 = h7(e13)) | (e21 = h7(e12)) | (e21 = h7(e11)) | (e21 = h7(e10))) & ((e20 = h7(e13)) | (e20 = h7(e12)) | (e20 = h7(e11)) | (e20 = h7(e10))) & (h7(op1(e13, e13)) = op2(h7(e13), h7(e13))) & (h7(op1(e13, e12)) = op2(h7(e13), h7(e12))) & (h7(op1(e13, e11)) = op2(h7(e13), h7(e11))) & (h7(op1(e13, e10)) = op2(h7(e13), h7(e10))) & (h7(op1(e12, e13)) = op2(h7(e12), h7(e13))) & (h7(op1(e12, e12)) = op2(h7(e12), h7(e12))) & (h7(op1(e12, e11)) = op2(h7(e12), h7(e11))) & (h7(op1(e12, e10)) = op2(h7(e12), h7(e10))) & (h7(op1(e11, e13)) = op2(h7(e11), h7(e13))) & (h7(op1(e11, e12)) = op2(h7(e11), h7(e12))) & (h7(op1(e11, e11)) = op2(h7(e11), h7(e11))) & (h7(op1(e11, e10)) = op2(h7(e11), h7(e10))) & (h7(op1(e10, e13)) = op2(h7(e10), h7(e13))) & (h7(op1(e10, e12)) = op2(h7(e10), h7(e12))) & (h7(op1(e10, e11)) = op2(h7(e10), h7(e11))) & (h7(op1(e10, e10)) = op2(h7(e10), h7(e10)))) | (((e23 = h6(e13)) | (e23 = h6(e12)) | (e23 = h6(e11)) | (e23 = h6(e10))) & ((e22 = h6(e13)) | (e22 = h6(e12)) | (e22 = h6(e11)) | (e22 = h6(e10))) & ((e21 = h6(e13)) | (e21 = h6(e12)) | (e21 = h6(e11)) | (e21 = h6(e10))) & ((e20 = h6(e13)) | (e20 = h6(e12)) | (e20 = h6(e11)) | (e20 = h6(e10))) & (h6(op1(e13, e13)) = op2(h6(e13), h6(e13))) & (h6(op1(e13, e12)) = op2(h6(e13), h6(e12))) & (h6(op1(e13, e11)) = op2(h6(e13), h6(e11))) & (h6(op1(e13, e10)) = op2(h6(e13), h6(e10))) & (h6(op1(e12, e13)) = op2(h6(e12), h6(e13))) & (h6(op1(e12, e12)) = op2(h6(e12), h6(e12))) & (h6(op1(e12, e11)) = op2(h6(e12), h6(e11))) & (h6(op1(e12, e10)) = op2(h6(e12), h6(e10))) & (h6(op1(e11, e13)) = op2(h6(e11), h6(e13))) & (h6(op1(e11, e12)) = op2(h6(e11), h6(e12))) & (h6(op1(e11, e11)) = op2(h6(e11), h6(e11))) & (h6(op1(e11, e10)) = op2(h6(e11), h6(e10))) & (h6(op1(e10, e13)) = op2(h6(e10), h6(e13))) & (h6(op1(e10, e12)) = op2(h6(e10), h6(e12))) & (h6(op1(e10, e11)) = op2(h6(e10), h6(e11))) & (h6(op1(e10, e10)) = op2(h6(e10), h6(e10)))) | (((e23 = h5(e13)) | (e23 = h5(e12)) | (e23 = h5(e11)) | (e23 = h5(e10))) & ((e22 = h5(e13)) | (e22 = h5(e12)) | (e22 = h5(e11)) | (e22 = h5(e10))) & ((e21 = h5(e13)) | (e21 = h5(e12)) | (e21 = h5(e11)) | (e21 = h5(e10))) & ((e20 = h5(e13)) | (e20 = h5(e12)) | (e20 = h5(e11)) | (e20 = h5(e10))) & (h5(op1(e13, e13)) = op2(h5(e13), h5(e13))) & (h5(op1(e13, e12)) = op2(h5(e13), h5(e12))) & (h5(op1(e13, e11)) = op2(h5(e13), h5(e11))) & (h5(op1(e13, e10)) = op2(h5(e13), h5(e10))) & (h5(op1(e12, e13)) = op2(h5(e12), h5(e13))) & (h5(op1(e12, e12)) = op2(h5(e12), h5(e12))) & (h5(op1(e12, e11)) = op2(h5(e12), h5(e11))) & (h5(op1(e12, e10)) = op2(h5(e12), h5(e10))) & (h5(op1(e11, e13)) = op2(h5(e11), h5(e13))) & (h5(op1(e11, e12)) = op2(h5(e11), h5(e12))) & (h5(op1(e11, e11)) = op2(h5(e11), h5(e11))) & (h5(op1(e11, e10)) = op2(h5(e11), h5(e10))) & (h5(op1(e10, e13)) = op2(h5(e10), h5(e13))) & (h5(op1(e10, e12)) = op2(h5(e10), h5(e12))) & (h5(op1(e10, e11)) = op2(h5(e10), h5(e11))) & (h5(op1(e10, e10)) = op2(h5(e10), h5(e10)))) | (((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', co1)).
fof(f1547, plain, (~ spl36_176 | ~ spl36_178), inference(avatar_split_clause, [], [f514, f1544, f1535])).
fof(f514, plain, (~ (e22 = h1(e10)) | ~ sP2), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP2), inference(nnf_transformation, [], [f31])).
fof(f1191, plain, spl36_81, inference(avatar_split_clause, [], [f332, f947])).
fof(f332, plain, (e20 = op2(e22, e23)), inference(cnf_transformation, [], [f13])).
fof(f1190, plain, spl36_17, inference(avatar_split_clause, [], [f330, f643])).
fof(f330, plain, (e10 = op1(e12, e13)), inference(cnf_transformation, [], [f12])).
fof(f1189, plain, spl36_125, inference(avatar_split_clause, [], [f326, f1134])).
fof(f326, plain, (e20 = op2(e20, e20)), inference(cnf_transformation, [], [f11])).
fof(f11, plain, ((e23 = op2(e23, e23)) & (e22 = op2(e22, e22)) & (e21 = op2(e21, e21)) & (e20 = op2(e20, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax11)).
fof(f1188, plain, spl36_106, inference(avatar_split_clause, [], [f327, f1053])).
fof(f327, plain, (e21 = op2(e21, e21)), inference(cnf_transformation, [], [f11])).
fof(f1187, plain, spl36_87, inference(avatar_split_clause, [], [f328, f972])).
fof(f328, plain, (e22 = op2(e22, e22)), inference(cnf_transformation, [], [f11])).
fof(f1186, plain, spl36_68, inference(avatar_split_clause, [], [f329, f891])).
fof(f329, plain, (e23 = op2(e23, e23)), inference(cnf_transformation, [], [f11])).
fof(f1185, plain, spl36_61, inference(avatar_split_clause, [], [f322, f830])).
fof(f322, plain, (e10 = op1(e10, e10)), inference(cnf_transformation, [], [f10])).
fof(f10, plain, ((e13 = op1(e13, e13)) & (e12 = op1(e12, e12)) & (e11 = op1(e11, e11)) & (e10 = op1(e10, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax10)).
fof(f1184, plain, spl36_42, inference(avatar_split_clause, [], [f323, f749])).
fof(f323, plain, (e11 = op1(e11, e11)), inference(cnf_transformation, [], [f10])).
fof(f1183, plain, spl36_23, inference(avatar_split_clause, [], [f324, f668])).
fof(f324, plain, (e12 = op1(e12, e12)), inference(cnf_transformation, [], [f10])).
fof(f1182, plain, spl36_4, inference(avatar_split_clause, [], [f325, f587])).
fof(f325, plain, (e13 = op1(e13, e13)), inference(cnf_transformation, [], [f10])).
fof(f1173, plain, (spl36_109 | spl36_105 | spl36_101 | spl36_97), inference(avatar_split_clause, [], [f174, f1015, f1032, f1049, f1066])).
fof(f174, plain, ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax4)).
fof(f1172, plain, (spl36_121 | spl36_105 | spl36_89 | spl36_73), inference(avatar_split_clause, [], [f175, f913, f981, f1049, f1117])).
fof(f175, plain, ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1159, plain, (spl36_96 | spl36_92 | spl36_88 | spl36_84), inference(avatar_split_clause, [], [f188, f959, f976, f993, f1010])).
fof(f188, plain, ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1155, plain, (spl36_78 | spl36_74 | spl36_70 | spl36_66), inference(avatar_split_clause, [], [f192, f883, f900, f917, f934])).
fof(f192, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1153, plain, (spl36_79 | spl36_75 | spl36_71 | spl36_67), inference(avatar_split_clause, [], [f194, f887, f904, f921, f938])).
fof(f194, plain, ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1152, plain, (spl36_115 | spl36_99 | spl36_83 | spl36_67), inference(avatar_split_clause, [], [f195, f887, f955, f1023, f1091])).
fof(f195, plain, ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1132, plain, (spl36_121 | spl36_122 | spl36_123 | spl36_124), inference(avatar_split_clause, [], [f151, f1129, f1125, f1121, f1117])).
fof(f151, plain, ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax3)).
fof(f1115, plain, (spl36_117 | spl36_118 | spl36_119 | spl36_120), inference(avatar_split_clause, [], [f152, f1112, f1108, f1104, f1100])).
fof(f152, plain, ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f3])).
fof(f1081, plain, (spl36_109 | spl36_110 | spl36_111 | spl36_112), inference(avatar_split_clause, [], [f154, f1078, f1074, f1070, f1066])).
fof(f154, plain, ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f3])).
fof(f1013, plain, (spl36_93 | spl36_94 | spl36_95 | spl36_96), inference(avatar_split_clause, [], [f158, f1010, f1006, f1002, f998])).
fof(f158, plain, ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))), inference(cnf_transformation, [], [f3])).
fof(f869, plain, (spl36_45 | spl36_41 | spl36_37 | spl36_33), inference(avatar_split_clause, [], [f126, f711, f728, f745, f762])).
fof(f126, plain, ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax2)).
fof(f868, plain, (spl36_57 | spl36_41 | spl36_25 | spl36_9), inference(avatar_split_clause, [], [f127, f609, f677, f745, f813])).
fof(f127, plain, ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f859, plain, (spl36_30 | spl36_26 | spl36_22 | spl36_18), inference(avatar_split_clause, [], [f136, f647, f664, f681, f698])).
fof(f136, plain, ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f858, plain, (spl36_54 | spl36_38 | spl36_22 | spl36_6), inference(avatar_split_clause, [], [f137, f596, f664, f732, f800])).
fof(f137, plain, ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f828, plain, (spl36_57 | spl36_58 | spl36_59 | spl36_60), inference(avatar_split_clause, [], [f103, f825, f821, f817, f813])).
fof(f103, plain, ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG113+1.p', ax1)).
fof(f811, plain, (spl36_53 | spl36_54 | spl36_55 | spl36_56), inference(avatar_split_clause, [], [f104, f808, f804, f800, f796])).
fof(f104, plain, ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f1])).
fof(f777, plain, (spl36_45 | spl36_46 | spl36_47 | spl36_48), inference(avatar_split_clause, [], [f106, f774, f770, f766, f762])).
fof(f106, plain, ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f1])).
fof(f726, plain, (spl36_33 | spl36_34 | spl36_35 | spl36_36), inference(avatar_split_clause, [], [f109, f723, f719, f715, f711])).
fof(f109, plain, ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))), inference(cnf_transformation, [], [f1])).
fof(f692, plain, (spl36_25 | spl36_26 | spl36_27 | spl36_28), inference(avatar_split_clause, [], [f111, f689, f685, f681, f677])).
fof(f111, plain, ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f641, plain, (spl36_13 | spl36_14 | spl36_15 | spl36_16), inference(avatar_split_clause, [], [f114, f638, f634, f630, f626])).
fof(f114, plain, ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f1])).