fof(f4472, plain, $false, inference(avatar_sat_refutation, [], [f377, f394, f411, f428, f445, f462, f479, f496, f513, f530, f547, f564, f581, f598, f615, f632, f633, f634, f635, f636, f637, f638, f639, f640, f641, f642, f643, f644, f645, f646, f647, f648, f649, f650, f651, f652, f653, f654, f655, f656, f657, f658, f659, f660, f661, f662, f663, f664, f670, f671, f672, f673, f674, f680, f681, f682, f683, f684, f690, f691, f692, f693, f694, f699, f701, f702, f703, f704, f710, f714, f719, f720, f721, f722, f723, f724, f729, f731, f732, f733, f734, f739, f740, f742, f743, f744, f749, f750, f752, f753, f754, f761, f764, f769, f770, f772, f773, f774, f779, f780, f781, f782, f783, f784, f789, f790, f791, f792, f793, f794, f799, f800, f801, f803, f804, f812, f814, f819, f820, f821, f826, f831, f836, f841, f846, f847, f848, f853, f858, f863, f868, f873, f874, f875, f880, f885, f890, f895, f896, f897, f899, f901, f902, f903, f904, f909, f914, f919, f924, f933, f942, f951, f960, f969, f978, f987, f996, f1001, f1006, f1011, f1016, f1021, f1026, f1031, f1036, f1037, f1038, f1039, f1040, f1041, f1042, f1043, f1044, f1054, f1074, f1080, f1085, f1111, f1126, f1129, f1139, f1147, f1167, f1170, f1173, f1177, f1182, f1189, f1192, f1194, f1197, f1200, f1203, f1218, f1228, f1233, f1240, f1244, f1255, f1270, f1287, f1289, f1300, f1334, f1338, f1345, f1353, f1365, f1371, f1381, f1385, f1389, f1399, f1405, f1411, f1418, f1451, f1455, f1469, f1473, f1505, f1512, f1515, f1529, f1549, f1563, f1588, f1602, f1615, f1644, f1647, f1660, f1665, f1687, f1706, f1707, f1714, f1732, f1735, f1742, f1744, f1752, f1763, f1781, f1784, f1803, f1815, f1839, f1847, f1848, f1859, f1861, f1882, f1904, f1938, f1947, f1948, f1951, f1964, f1988, f1991, f1996, f2000, f2003, f2009, f2011, f2033, f2039, f2070, f2076, f2113, f2128, f2139, f2158, f2193, f2203, f2228, f2230, f2264, f2268, f2274, f2293, f2298, f2318, f2322, f2334, f2338, f2363, f2366, f2402, f2427, f2431, f2434, f2436, f2462, f2464, f2468, f2494, f2495, f2533, f2545, f2550, f2570, f2575, f2582, f2585, f2592, f2600, f2602, f2604, f2620, f2621, f2623, f2647, f2648, f2649, f2655, f2658, f2689, f2706, f2712, f2716, f2732, f2746, f2748, f2756, f2757, f2761, f2774, f2782, f2823, f2845, f2847, f2849, f2852, f2855, f2874, f2890, f2891, f2907, f2910, f2914, f2920, f2922, f2927, f2933, f2938, f2968, f2973, f2976, f2981, f3001, f3006, f3007, f3008, f3011, f3013, f3038, f3046, f3079, f3089, f3102, f3103, f3108, f3140, f3145, f3166, f3170, f3173, f3199, f3211, f3212, f3219, f3222, f3233, f3234, f3241, f3248, f3249, f3263, f3271, f3287, f3301, f3318, f3354, f3362, f3363, f3374, f3379, f3389, f3390, f3425, f3426, f3428, f3433, f3434, f3448, f3454, f3455, f3499, f3502, f3510, f3525, f3528, f3534, f3536, f3538, f3540, f3542, f3545, f3554, f3560, f3563, f3571, f3581, f3582, f3586, f3591, f3601, f3617, f3618, f3621, f3645, f3652, f3654, f3668, f3686, f3687, f3702, f3709, f3714, f3730, f3735, f3745, f3763, f3775, f3779, f3808, f3810, f3816, f3845, f3846, f3856, f3869, f3878, f3891, f3892, f3893, f3900, f3915, f3925, f3949, f3962, f3963, f3964, f3965, f3973, f3974, f3978, f3986, f3987, f3989, f3994, f3995, f4000, f4006, f4007, f4021, f4022, f4036, f4041, f4062, f4074, f4075, f4076, f4088, f4096, f4102, f4105, f4114, f4126, f4127, f4129, f4132, f4152, f4164, f4195, f4197, f4221, f4225, f4228, f4237, f4238, f4244, f4251, f4252, f4257, f4262, f4276, f4308, f4321, f4328, f4330, f4336, f4337, f4338, f4343, f4349, f4357, f4358, f4363, f4368, f4372, f4374, f4383, f4385, f4387, f4389, f4396, f4403, f4405, f4408, f4419, f4420, f4423, f4427, f4435, f4436, f4438, f4446, f4449, f4453, f4460])).
fof(f4460, plain, (~ spl18_1 | ~ spl18_5), inference(avatar_split_clause, [], [f4455, f379, f362])).
fof(f362, plain, (spl18_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_1])])).
fof(f379, plain, (spl18_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_5])])).
fof(f4455, plain, (~ (e0 = op(e3, e3)) | ~ spl18_5), inference(backward_demodulation, [], [f186, f381])).
fof(f381, plain, ((e0 = op(e3, e2)) | ~ spl18_5), inference(avatar_component_clause, [], [f379])).
fof(f186, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax3)).
fof(f4453, plain, (~ spl18_8 | ~ spl18_16), inference(avatar_split_clause, [], [f4450, f425, f391])).
fof(f391, plain, (spl18_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_8])])).
fof(f425, plain, (spl18_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_16])])).
fof(f4450, plain, (~ (e3 = op(e3, e2)) | ~ spl18_16), inference(backward_demodulation, [], [f182, f427])).
fof(f427, plain, ((e3 = op(e3, e0)) | ~ spl18_16), inference(avatar_component_clause, [], [f425])).
fof(f182, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4449, plain, (~ spl18_17 | ~ spl18_31 | ~ spl18_81), inference(avatar_contradiction_clause, [], [f4448])).
fof(f4448, plain, ($false | (~ spl18_17 | ~ spl18_31 | ~ spl18_81)), inference(subsumption_resolution, [], [f4447, f192])).
fof(f192, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax4)).
fof(f4447, plain, ((e2 = e3) | (~ spl18_17 | ~ spl18_31 | ~ spl18_81)), inference(forward_demodulation, [], [f4445, f491])).
fof(f491, plain, ((e2 = op(e2, e0)) | ~ spl18_31), inference(avatar_component_clause, [], [f489])).
fof(f489, plain, (spl18_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_31])])).
fof(f4445, plain, ((e3 = op(e2, e0)) | (~ spl18_17 | ~ spl18_81)), inference(backward_demodulation, [], [f825, f432])).
fof(f432, plain, ((e0 = op(e2, e3)) | ~ spl18_17), inference(avatar_component_clause, [], [f430])).
fof(f430, plain, (spl18_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_17])])).
fof(f825, plain, ((e3 = op(e2, op(e2, e3))) | ~ spl18_81), inference(avatar_component_clause, [], [f823])).
fof(f823, plain, (spl18_81 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl18_81])])).
fof(f4446, plain, (~ spl18_1 | ~ spl18_17), inference(avatar_split_clause, [], [f4441, f430, f362])).
fof(f4441, plain, (~ (e0 = op(e3, e3)) | ~ spl18_17), inference(backward_demodulation, [], [f162, f432])).
fof(f162, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4438, plain, (~ spl18_1 | ~ spl18_24 | spl18_116), inference(avatar_split_clause, [], [f4432, f1003, f459, f362])).
fof(f459, plain, (spl18_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_24])])).
fof(f1003, plain, (spl18_116 <=> (e0 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl18_116])])).
fof(f4432, plain, (~ (e0 = op(e3, e3)) | (~ spl18_24 | spl18_116)), inference(backward_demodulation, [], [f1005, f461])).
fof(f461, plain, ((e3 = op(e2, e2)) | ~ spl18_24), inference(avatar_component_clause, [], [f459])).
fof(f1005, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | spl18_116), inference(avatar_component_clause, [], [f1003])).
fof(f4436, plain, (~ spl18_20 | ~ spl18_24), inference(avatar_split_clause, [], [f4429, f459, f442])).
fof(f442, plain, (spl18_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_20])])).
fof(f4429, plain, (~ (e3 = op(e2, e3)) | ~ spl18_24), inference(backward_demodulation, [], [f180, f461])).
fof(f180, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4435, plain, (~ spl18_8 | ~ spl18_24), inference(avatar_split_clause, [], [f4428, f459, f391])).
fof(f4428, plain, (~ (e3 = op(e3, e2)) | ~ spl18_24), inference(backward_demodulation, [], [f156, f461])).
fof(f156, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4427, plain, (~ spl18_20 | ~ spl18_52), inference(avatar_split_clause, [], [f4424, f578, f442])).
fof(f578, plain, (spl18_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_52])])).
fof(f4424, plain, (~ (e3 = op(e2, e3)) | ~ spl18_52), inference(backward_demodulation, [], [f158, f580])).
fof(f580, plain, ((e3 = op(e0, e3)) | ~ spl18_52), inference(avatar_component_clause, [], [f578])).
fof(f158, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4423, plain, (~ spl18_11 | ~ spl18_62 | spl18_111 | ~ spl18_112), inference(avatar_contradiction_clause, [], [f4422])).
fof(f4422, plain, ($false | (~ spl18_11 | ~ spl18_62 | spl18_111 | ~ spl18_112)), inference(subsumption_resolution, [], [f4418, f406])).
fof(f406, plain, ((e2 = op(e3, e1)) | ~ spl18_11), inference(avatar_component_clause, [], [f404])).
fof(f404, plain, (spl18_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_11])])).
fof(f4418, plain, (~ (e2 = op(e3, e1)) | (~ spl18_62 | spl18_111 | ~ spl18_112)), inference(backward_demodulation, [], [f4400, f623])).
fof(f623, plain, ((op(e0, e0) = e1) | ~ spl18_62), inference(avatar_component_clause, [], [f621])).
fof(f621, plain, (spl18_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl18_62])])).
fof(f4400, plain, (~ (e2 = op(e3, op(e0, e0))) | (spl18_111 | ~ spl18_112)), inference(backward_demodulation, [], [f982, f985])).
fof(f985, plain, ((e3 = op(op(e0, e0), op(e0, e0))) | ~ spl18_112), inference(avatar_component_clause, [], [f984])).
fof(f984, plain, (spl18_112 <=> (e3 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl18_112])])).
fof(f982, plain, (~ (e2 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | spl18_111), inference(avatar_component_clause, [], [f980])).
fof(f980, plain, (spl18_111 <=> (e2 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl18_111])])).
fof(f4420, plain, (~ spl18_50 | ~ spl18_62), inference(avatar_split_clause, [], [f4410, f621, f570])).
fof(f570, plain, (spl18_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_50])])).
fof(f4410, plain, (~ (e1 = op(e0, e3)) | ~ spl18_62), inference(backward_demodulation, [], [f165, f623])).
fof(f165, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4419, plain, (~ spl18_14 | ~ spl18_62), inference(avatar_split_clause, [], [f4409, f621, f417])).
fof(f417, plain, (spl18_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_14])])).
fof(f4409, plain, (~ (e1 = op(e3, e0)) | ~ spl18_62), inference(backward_demodulation, [], [f141, f623])).
fof(f141, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f4408, plain, (~ spl18_21 | ~ spl18_31 | spl18_84), inference(avatar_split_clause, [], [f4407, f838, f489, f447])).
fof(f447, plain, (spl18_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_21])])).
fof(f838, plain, (spl18_84 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl18_84])])).
fof(f4407, plain, (~ (e0 = op(e2, e2)) | (~ spl18_31 | spl18_84)), inference(forward_demodulation, [], [f839, f491])).
fof(f839, plain, (~ (e0 = op(e2, op(e2, e0))) | spl18_84), inference(avatar_component_clause, [], [f838])).
fof(f4405, plain, (spl18_62 | ~ spl18_57 | ~ spl18_93), inference(avatar_split_clause, [], [f4404, f887, f600, f621])).
fof(f600, plain, (spl18_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_57])])).
fof(f887, plain, (spl18_93 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl18_93])])).
fof(f4404, plain, ((op(e0, e0) = e1) | (~ spl18_57 | ~ spl18_93)), inference(forward_demodulation, [], [f889, f602])).
fof(f602, plain, ((e0 = op(e0, e1)) | ~ spl18_57), inference(avatar_component_clause, [], [f600])).
fof(f889, plain, ((e1 = op(e0, op(e0, e1))) | ~ spl18_93), inference(avatar_component_clause, [], [f887])).
fof(f4403, plain, (~ spl18_104 | ~ spl18_106), inference(avatar_contradiction_clause, [], [f4402])).
fof(f4402, plain, ($false | (~ spl18_104 | ~ spl18_106)), inference(subsumption_resolution, [], [f4401, f191])).
fof(f191, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f4401, plain, ((e1 = e3) | (~ spl18_104 | ~ spl18_106)), inference(backward_demodulation, [], [f949, f958])).
fof(f958, plain, ((e1 = op(op(e2, e2), op(e2, e2))) | ~ spl18_106), inference(avatar_component_clause, [], [f957])).
fof(f957, plain, (spl18_106 <=> (e1 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl18_106])])).
fof(f949, plain, ((e3 = op(op(e2, e2), op(e2, e2))) | ~ spl18_104), inference(avatar_component_clause, [], [f948])).
fof(f948, plain, (spl18_104 <=> (e3 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl18_104])])).
fof(f4396, plain, (~ spl18_1 | ~ spl18_44 | spl18_122), inference(avatar_split_clause, [], [f4395, f1033, f544, f362])).
fof(f544, plain, (spl18_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_44])])).
fof(f1033, plain, (spl18_122 <=> (e0 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl18_122])])).
fof(f4395, plain, (~ (e0 = op(e3, e3)) | (~ spl18_44 | spl18_122)), inference(forward_demodulation, [], [f1035, f546])).
fof(f546, plain, ((e3 = op(e1, e1)) | ~ spl18_44), inference(avatar_component_clause, [], [f544])).
fof(f1035, plain, (~ (e0 = op(op(e1, e1), op(e1, e1))) | spl18_122), inference(avatar_component_clause, [], [f1033])).
fof(f4389, plain, (~ spl18_22 | ~ spl18_26), inference(avatar_split_clause, [], [f4388, f468, f451])).
fof(f451, plain, (spl18_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_22])])).
fof(f468, plain, (spl18_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_26])])).
fof(f4388, plain, (~ (e1 = op(e2, e2)) | ~ spl18_26), inference(forward_demodulation, [], [f178, f470])).
fof(f470, plain, ((e1 = op(e2, e1)) | ~ spl18_26), inference(avatar_component_clause, [], [f468])).
fof(f178, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4387, plain, (~ spl18_23 | ~ spl18_31), inference(avatar_split_clause, [], [f4386, f489, f455])).
fof(f455, plain, (spl18_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_23])])).
fof(f4386, plain, (~ (e2 = op(e2, e2)) | ~ spl18_31), inference(forward_demodulation, [], [f176, f491])).
fof(f176, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4385, plain, (~ spl18_22 | ~ spl18_38), inference(avatar_split_clause, [], [f4384, f519, f451])).
fof(f519, plain, (spl18_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_38])])).
fof(f4384, plain, (~ (e1 = op(e2, e2)) | ~ spl18_38), inference(forward_demodulation, [], [f154, f521])).
fof(f521, plain, ((e1 = op(e1, e2)) | ~ spl18_38), inference(avatar_component_clause, [], [f519])).
fof(f154, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4383, plain, (~ spl18_23 | ~ spl18_55), inference(avatar_split_clause, [], [f4382, f591, f455])).
fof(f591, plain, (spl18_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_55])])).
fof(f4382, plain, (~ (e2 = op(e2, e2)) | ~ spl18_55), inference(forward_demodulation, [], [f152, f593])).
fof(f593, plain, ((e2 = op(e0, e2)) | ~ spl18_55), inference(avatar_component_clause, [], [f591])).
fof(f152, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4374, plain, (~ spl18_3 | ~ spl18_11), inference(avatar_split_clause, [], [f4373, f404, f370])).
fof(f370, plain, (spl18_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_3])])).
fof(f4373, plain, (~ (e2 = op(e3, e3)) | ~ spl18_11), inference(forward_demodulation, [], [f185, f406])).
fof(f185, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4372, plain, (~ spl18_3 | ~ spl18_35), inference(avatar_split_clause, [], [f4371, f506, f370])).
fof(f506, plain, (spl18_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_35])])).
fof(f4371, plain, (~ (e2 = op(e3, e3)) | ~ spl18_35), inference(forward_demodulation, [], [f161, f508])).
fof(f508, plain, ((e2 = op(e1, e3)) | ~ spl18_35), inference(avatar_component_clause, [], [f506])).
fof(f161, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4368, plain, (~ spl18_10 | ~ spl18_14), inference(avatar_split_clause, [], [f4364, f417, f400])).
fof(f400, plain, (spl18_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_10])])).
fof(f4364, plain, (~ (e1 = op(e3, e1)) | ~ spl18_14), inference(backward_demodulation, [], [f181, f419])).
fof(f419, plain, ((e1 = op(e3, e0)) | ~ spl18_14), inference(avatar_component_clause, [], [f417])).
fof(f181, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4363, plain, (~ spl18_10 | ~ spl18_26), inference(avatar_split_clause, [], [f4359, f468, f400])).
fof(f4359, plain, (~ (e1 = op(e3, e1)) | ~ spl18_26), inference(backward_demodulation, [], [f150, f470])).
fof(f150, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4358, plain, (~ spl18_15 | ~ spl18_31), inference(avatar_split_clause, [], [f4356, f489, f421])).
fof(f421, plain, (spl18_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_15])])).
fof(f4356, plain, (~ (e2 = op(e3, e0)) | ~ spl18_31), inference(backward_demodulation, [], [f144, f491])).
fof(f144, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f4357, plain, (~ spl18_27 | ~ spl18_31), inference(avatar_split_clause, [], [f4355, f489, f472])).
fof(f472, plain, (spl18_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_27])])).
fof(f4355, plain, (~ (e2 = op(e2, e1)) | ~ spl18_31), inference(backward_demodulation, [], [f175, f491])).
fof(f175, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4349, plain, (~ spl18_34 | ~ spl18_38), inference(avatar_split_clause, [], [f4345, f519, f502])).
fof(f502, plain, (spl18_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_34])])).
fof(f4345, plain, (~ (e1 = op(e1, e3)) | ~ spl18_38), inference(backward_demodulation, [], [f174, f521])).
fof(f174, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4343, plain, (~ spl18_34 | ~ spl18_50), inference(avatar_split_clause, [], [f4339, f570, f502])).
fof(f4339, plain, (~ (e1 = op(e1, e3)) | ~ spl18_50), inference(backward_demodulation, [], [f157, f572])).
fof(f572, plain, ((e1 = op(e0, e3)) | ~ spl18_50), inference(avatar_component_clause, [], [f570])).
fof(f157, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4338, plain, (~ spl18_39 | ~ spl18_55), inference(avatar_split_clause, [], [f4334, f591, f523])).
fof(f523, plain, (spl18_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_39])])).
fof(f4334, plain, (~ (e2 = op(e1, e2)) | ~ spl18_55), inference(backward_demodulation, [], [f151, f593])).
fof(f151, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f4337, plain, (~ spl18_51 | ~ spl18_55), inference(avatar_split_clause, [], [f4333, f591, f574])).
fof(f574, plain, (spl18_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_51])])).
fof(f4333, plain, (~ (e2 = op(e0, e3)) | ~ spl18_55), inference(backward_demodulation, [], [f168, f593])).
fof(f168, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f4336, plain, (~ spl18_55 | spl18_92), inference(avatar_contradiction_clause, [], [f4335])).
fof(f4335, plain, ($false | (~ spl18_55 | spl18_92)), inference(subsumption_resolution, [], [f4332, f593])).
fof(f4332, plain, (~ (e2 = op(e0, e2)) | (~ spl18_55 | spl18_92)), inference(backward_demodulation, [], [f883, f593])).
fof(f883, plain, (~ (e2 = op(e0, op(e0, e2))) | spl18_92), inference(avatar_component_clause, [], [f882])).
fof(f882, plain, (spl18_92 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl18_92])])).
fof(f4330, plain, (spl18_31 | ~ spl18_21 | ~ spl18_82), inference(avatar_split_clause, [], [f4329, f828, f447, f489])).
fof(f828, plain, (spl18_82 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl18_82])])).
fof(f4329, plain, ((e2 = op(e2, e0)) | (~ spl18_21 | ~ spl18_82)), inference(forward_demodulation, [], [f830, f449])).
fof(f449, plain, ((e0 = op(e2, e2)) | ~ spl18_21), inference(avatar_component_clause, [], [f447])).
fof(f830, plain, ((e2 = op(e2, op(e2, e2))) | ~ spl18_82), inference(avatar_component_clause, [], [f828])).
fof(f4328, plain, (~ spl18_34 | ~ spl18_44 | spl18_88), inference(avatar_split_clause, [], [f4327, f860, f544, f502])).
fof(f860, plain, (spl18_88 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl18_88])])).
fof(f4327, plain, (~ (e1 = op(e1, e3)) | (~ spl18_44 | spl18_88)), inference(forward_demodulation, [], [f861, f546])).
fof(f861, plain, (~ (e1 = op(e1, op(e1, e1))) | spl18_88), inference(avatar_component_clause, [], [f860])).
fof(f4321, plain, (~ spl18_51 | ~ spl18_1 | ~ spl18_44 | spl18_99), inference(avatar_split_clause, [], [f4300, f926, f544, f362, f574])).
fof(f926, plain, (spl18_99 <=> (e2 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl18_99])])).
fof(f4300, plain, (~ (e2 = op(e0, e3)) | (~ spl18_1 | ~ spl18_44 | spl18_99)), inference(forward_demodulation, [], [f4299, f364])).
fof(f364, plain, ((e0 = op(e3, e3)) | ~ spl18_1), inference(avatar_component_clause, [], [f362])).
fof(f4299, plain, (~ (e2 = op(op(e3, e3), e3)) | (~ spl18_44 | spl18_99)), inference(forward_demodulation, [], [f928, f546])).
fof(f928, plain, (~ (e2 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | spl18_99), inference(avatar_component_clause, [], [f926])).
fof(f4308, plain, (~ spl18_14 | ~ spl18_21 | spl18_103 | ~ spl18_104), inference(avatar_split_clause, [], [f4270, f948, f944, f447, f417])).
fof(f944, plain, (spl18_103 <=> (e1 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl18_103])])).
fof(f4270, plain, (~ (e1 = op(e3, e0)) | (~ spl18_21 | spl18_103 | ~ spl18_104)), inference(backward_demodulation, [], [f4226, f449])).
fof(f4226, plain, (~ (e1 = op(e3, op(e2, e2))) | (spl18_103 | ~ spl18_104)), inference(forward_demodulation, [], [f946, f949])).
fof(f946, plain, (~ (e1 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | spl18_103), inference(avatar_component_clause, [], [f944])).
fof(f4276, plain, (~ spl18_2 | ~ spl18_10), inference(avatar_split_clause, [], [f4273, f400, f366])).
fof(f366, plain, (spl18_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_2])])).
fof(f4273, plain, (~ (e1 = op(e3, e3)) | ~ spl18_10), inference(backward_demodulation, [], [f185, f402])).
fof(f402, plain, ((e1 = op(e3, e1)) | ~ spl18_10), inference(avatar_component_clause, [], [f400])).
fof(f4262, plain, (~ spl18_22 | ~ spl18_30), inference(avatar_split_clause, [], [f4261, f485, f451])).
fof(f485, plain, (spl18_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_30])])).
fof(f4261, plain, (~ (e1 = op(e2, e2)) | ~ spl18_30), inference(backward_demodulation, [], [f176, f487])).
fof(f487, plain, ((e1 = op(e2, e0)) | ~ spl18_30), inference(avatar_component_clause, [], [f485])).
fof(f4257, plain, (~ spl18_2 | ~ spl18_34), inference(avatar_split_clause, [], [f4254, f502, f366])).
fof(f4254, plain, (~ (e1 = op(e3, e3)) | ~ spl18_34), inference(backward_demodulation, [], [f161, f504])).
fof(f504, plain, ((e1 = op(e1, e3)) | ~ spl18_34), inference(avatar_component_clause, [], [f502])).
fof(f4252, plain, (~ spl18_33 | ~ spl18_45), inference(avatar_split_clause, [], [f4248, f549, f498])).
fof(f498, plain, (spl18_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_33])])).
fof(f549, plain, (spl18_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_45])])).
fof(f4248, plain, (~ (e0 = op(e1, e3)) | ~ spl18_45), inference(backward_demodulation, [], [f171, f551])).
fof(f551, plain, ((e0 = op(e1, e0)) | ~ spl18_45), inference(avatar_component_clause, [], [f549])).
fof(f171, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f4251, plain, (~ spl18_29 | ~ spl18_45), inference(avatar_split_clause, [], [f4247, f549, f481])).
fof(f481, plain, (spl18_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_29])])).
fof(f4247, plain, (~ (e0 = op(e2, e0)) | ~ spl18_45), inference(backward_demodulation, [], [f142, f551])).
fof(f142, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f4244, plain, (~ spl18_22 | ~ spl18_54), inference(avatar_split_clause, [], [f4242, f587, f451])).
fof(f587, plain, (spl18_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_54])])).
fof(f4242, plain, (~ (e1 = op(e2, e2)) | ~ spl18_54), inference(backward_demodulation, [], [f152, f589])).
fof(f589, plain, ((e1 = op(e0, e2)) | ~ spl18_54), inference(avatar_component_clause, [], [f587])).
fof(f4238, plain, (~ spl18_53 | ~ spl18_57), inference(avatar_split_clause, [], [f4234, f600, f583])).
fof(f583, plain, (spl18_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_53])])).
fof(f4234, plain, (~ (e0 = op(e0, e2)) | ~ spl18_57), inference(backward_demodulation, [], [f166, f602])).
fof(f166, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f4237, plain, (~ spl18_9 | ~ spl18_57), inference(avatar_split_clause, [], [f4233, f600, f396])).
fof(f396, plain, (spl18_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_9])])).
fof(f4233, plain, (~ (e0 = op(e3, e1)) | ~ spl18_57), inference(backward_demodulation, [], [f147, f602])).
fof(f147, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4228, plain, (spl18_34 | ~ spl18_44 | ~ spl18_88), inference(avatar_split_clause, [], [f4227, f860, f544, f502])).
fof(f4227, plain, ((e1 = op(e1, e3)) | (~ spl18_44 | ~ spl18_88)), inference(forward_demodulation, [], [f862, f546])).
fof(f862, plain, ((e1 = op(e1, op(e1, e1))) | ~ spl18_88), inference(avatar_component_clause, [], [f860])).
fof(f4225, plain, (~ spl18_2 | ~ spl18_64 | spl18_120), inference(avatar_split_clause, [], [f4224, f1023, f629, f366])).
fof(f629, plain, (spl18_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl18_64])])).
fof(f1023, plain, (spl18_120 <=> (e1 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl18_120])])).
fof(f4224, plain, (~ (e1 = op(e3, e3)) | (~ spl18_64 | spl18_120)), inference(forward_demodulation, [], [f1025, f631])).
fof(f631, plain, ((op(e0, e0) = e3) | ~ spl18_64), inference(avatar_component_clause, [], [f629])).
fof(f1025, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | spl18_120), inference(avatar_component_clause, [], [f1023])).
fof(f4221, plain, (spl18_3 | ~ spl18_64 | ~ spl18_114), inference(avatar_contradiction_clause, [], [f4220])).
fof(f4220, plain, ($false | (spl18_3 | ~ spl18_64 | ~ spl18_114)), inference(subsumption_resolution, [], [f4219, f371])).
fof(f371, plain, (~ (e2 = op(e3, e3)) | spl18_3), inference(avatar_component_clause, [], [f370])).
fof(f4219, plain, ((e2 = op(e3, e3)) | (~ spl18_64 | ~ spl18_114)), inference(forward_demodulation, [], [f994, f631])).
fof(f994, plain, ((e2 = op(op(e0, e0), op(e0, e0))) | ~ spl18_114), inference(avatar_component_clause, [], [f993])).
fof(f993, plain, (spl18_114 <=> (e2 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl18_114])])).
fof(f4197, plain, (~ spl18_23 | ~ spl18_27), inference(avatar_split_clause, [], [f4196, f472, f455])).
fof(f4196, plain, (~ (e2 = op(e2, e2)) | ~ spl18_27), inference(forward_demodulation, [], [f178, f474])).
fof(f474, plain, ((e2 = op(e2, e1)) | ~ spl18_27), inference(avatar_component_clause, [], [f472])).
fof(f4195, plain, (~ spl18_23 | ~ spl18_39), inference(avatar_split_clause, [], [f4194, f523, f455])).
fof(f4194, plain, (~ (e2 = op(e2, e2)) | ~ spl18_39), inference(forward_demodulation, [], [f154, f525])).
fof(f525, plain, ((e2 = op(e1, e2)) | ~ spl18_39), inference(avatar_component_clause, [], [f523])).
fof(f4164, plain, (~ spl18_3 | ~ spl18_15), inference(avatar_split_clause, [], [f4161, f421, f370])).
fof(f4161, plain, (~ (e2 = op(e3, e3)) | ~ spl18_15), inference(backward_demodulation, [], [f183, f423])).
fof(f423, plain, ((e2 = op(e3, e0)) | ~ spl18_15), inference(avatar_component_clause, [], [f421])).
fof(f183, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4152, plain, (~ spl18_17 | ~ spl18_29), inference(avatar_split_clause, [], [f4147, f481, f430])).
fof(f4147, plain, (~ (e0 = op(e2, e3)) | ~ spl18_29), inference(backward_demodulation, [], [f177, f483])).
fof(f483, plain, ((e0 = op(e2, e0)) | ~ spl18_29), inference(avatar_component_clause, [], [f481])).
fof(f177, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f4132, plain, (~ spl18_9 | ~ spl18_22 | ~ spl18_44 | spl18_115), inference(avatar_split_clause, [], [f4125, f998, f544, f451, f396])).
fof(f998, plain, (spl18_115 <=> (e0 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl18_115])])).
fof(f4125, plain, (~ (e0 = op(e3, e1)) | (~ spl18_22 | ~ spl18_44 | spl18_115)), inference(backward_demodulation, [], [f4057, f546])).
fof(f4057, plain, (~ (e0 = op(op(e1, e1), e1)) | (~ spl18_22 | spl18_115)), inference(forward_demodulation, [], [f1000, f453])).
fof(f453, plain, ((e1 = op(e2, e2)) | ~ spl18_22), inference(avatar_component_clause, [], [f451])).
fof(f1000, plain, (~ (e0 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | spl18_115), inference(avatar_component_clause, [], [f998])).
fof(f4129, plain, (~ spl18_3 | ~ spl18_44 | spl18_102), inference(avatar_split_clause, [], [f4121, f939, f544, f370])).
fof(f939, plain, (spl18_102 <=> (e2 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl18_102])])).
fof(f4121, plain, (~ (e2 = op(e3, e3)) | (~ spl18_44 | spl18_102)), inference(backward_demodulation, [], [f941, f546])).
fof(f941, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | spl18_102), inference(avatar_component_clause, [], [f939])).
fof(f4127, plain, (~ spl18_40 | ~ spl18_44), inference(avatar_split_clause, [], [f4117, f544, f527])).
fof(f527, plain, (spl18_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_40])])).
fof(f4117, plain, (~ (e3 = op(e1, e2)) | ~ spl18_44), inference(backward_demodulation, [], [f172, f546])).
fof(f172, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f4126, plain, (~ spl18_28 | ~ spl18_44), inference(avatar_split_clause, [], [f4115, f544, f476])).
fof(f476, plain, (spl18_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_28])])).
fof(f4115, plain, (~ (e3 = op(e2, e1)) | ~ spl18_44), inference(backward_demodulation, [], [f148, f546])).
fof(f148, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f4114, plain, (~ spl18_34 | ~ spl18_46), inference(avatar_split_clause, [], [f4110, f553, f502])).
fof(f553, plain, (spl18_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_46])])).
fof(f4110, plain, (~ (e1 = op(e1, e3)) | ~ spl18_46), inference(backward_demodulation, [], [f171, f555])).
fof(f555, plain, ((e1 = op(e1, e0)) | ~ spl18_46), inference(avatar_component_clause, [], [f553])).
fof(f4105, plain, (~ spl18_51 | ~ spl18_53 | ~ spl18_91), inference(avatar_contradiction_clause, [], [f4104])).
fof(f4104, plain, ($false | (~ spl18_51 | ~ spl18_53 | ~ spl18_91)), inference(subsumption_resolution, [], [f4103, f189])).
fof(f189, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f4103, plain, ((e0 = e3) | (~ spl18_51 | ~ spl18_53 | ~ spl18_91)), inference(forward_demodulation, [], [f4101, f585])).
fof(f585, plain, ((e0 = op(e0, e2)) | ~ spl18_53), inference(avatar_component_clause, [], [f583])).
fof(f4101, plain, ((e3 = op(e0, e2)) | (~ spl18_51 | ~ spl18_91)), inference(backward_demodulation, [], [f879, f576])).
fof(f576, plain, ((e2 = op(e0, e3)) | ~ spl18_51), inference(avatar_component_clause, [], [f574])).
fof(f879, plain, ((e3 = op(e0, op(e0, e3))) | ~ spl18_91), inference(avatar_component_clause, [], [f877])).
fof(f877, plain, (spl18_91 <=> (e3 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl18_91])])).
fof(f4102, plain, (~ spl18_3 | ~ spl18_51), inference(avatar_split_clause, [], [f4099, f574, f370])).
fof(f4099, plain, (~ (e2 = op(e3, e3)) | ~ spl18_51), inference(backward_demodulation, [], [f159, f576])).
fof(f159, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4096, plain, (~ spl18_5 | ~ spl18_53), inference(avatar_split_clause, [], [f4090, f583, f379])).
fof(f4090, plain, (~ (e0 = op(e3, e2)) | ~ spl18_53), inference(backward_demodulation, [], [f153, f585])).
fof(f153, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4088, plain, (~ spl18_10 | ~ spl18_58), inference(avatar_split_clause, [], [f4082, f604, f400])).
fof(f604, plain, (spl18_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_58])])).
fof(f4082, plain, (~ (e1 = op(e3, e1)) | ~ spl18_58), inference(backward_demodulation, [], [f147, f606])).
fof(f606, plain, ((e1 = op(e0, e1)) | ~ spl18_58), inference(avatar_component_clause, [], [f604])).
fof(f4076, plain, (spl18_2 | ~ spl18_64 | ~ spl18_120), inference(avatar_split_clause, [], [f4070, f1023, f629, f366])).
fof(f4070, plain, ((e1 = op(e3, e3)) | (~ spl18_64 | ~ spl18_120)), inference(backward_demodulation, [], [f1024, f631])).
fof(f1024, plain, ((e1 = op(op(e0, e0), op(e0, e0))) | ~ spl18_120), inference(avatar_component_clause, [], [f1023])).
fof(f4075, plain, (~ spl18_52 | ~ spl18_64), inference(avatar_split_clause, [], [f4068, f629, f578])).
fof(f4068, plain, (~ (e3 = op(e0, e3)) | ~ spl18_64), inference(backward_demodulation, [], [f165, f631])).
fof(f4074, plain, (~ spl18_16 | ~ spl18_64), inference(avatar_split_clause, [], [f4065, f629, f425])).
fof(f4065, plain, (~ (e3 = op(e3, e0)) | ~ spl18_64), inference(backward_demodulation, [], [f141, f631])).
fof(f4062, plain, (spl18_27 | ~ spl18_22 | ~ spl18_82), inference(avatar_split_clause, [], [f4061, f828, f451, f472])).
fof(f4061, plain, ((e2 = op(e2, e1)) | (~ spl18_22 | ~ spl18_82)), inference(forward_demodulation, [], [f830, f453])).
fof(f4041, plain, (~ spl18_6 | ~ spl18_10), inference(avatar_split_clause, [], [f4037, f400, f383])).
fof(f383, plain, (spl18_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_6])])).
fof(f4037, plain, (~ (e1 = op(e3, e2)) | ~ spl18_10), inference(backward_demodulation, [], [f184, f402])).
fof(f184, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4036, plain, (~ spl18_12 | ~ spl18_16), inference(avatar_split_clause, [], [f4033, f425, f408])).
fof(f408, plain, (spl18_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_12])])).
fof(f4033, plain, (~ (e3 = op(e3, e1)) | ~ spl18_16), inference(backward_demodulation, [], [f181, f427])).
fof(f4022, plain, (~ spl18_18 | ~ spl18_22), inference(avatar_split_clause, [], [f4010, f451, f434])).
fof(f434, plain, (spl18_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_18])])).
fof(f4010, plain, (~ (e1 = op(e2, e3)) | ~ spl18_22), inference(backward_demodulation, [], [f180, f453])).
fof(f4021, plain, (~ spl18_6 | ~ spl18_22), inference(avatar_split_clause, [], [f4009, f451, f383])).
fof(f4009, plain, (~ (e1 = op(e3, e2)) | ~ spl18_22), inference(backward_demodulation, [], [f156, f453])).
fof(f4007, plain, (~ spl18_24 | ~ spl18_28), inference(avatar_split_clause, [], [f4003, f476, f459])).
fof(f4003, plain, (~ (e3 = op(e2, e2)) | ~ spl18_28), inference(backward_demodulation, [], [f178, f478])).
fof(f478, plain, ((e3 = op(e2, e1)) | ~ spl18_28), inference(avatar_component_clause, [], [f476])).
fof(f4006, plain, (~ spl18_12 | ~ spl18_28), inference(avatar_split_clause, [], [f4002, f476, f408])).
fof(f4002, plain, (~ (e3 = op(e3, e1)) | ~ spl18_28), inference(backward_demodulation, [], [f150, f478])).
fof(f4000, plain, (~ spl18_18 | ~ spl18_34), inference(avatar_split_clause, [], [f3996, f502, f434])).
fof(f3996, plain, (~ (e1 = op(e2, e3)) | ~ spl18_34), inference(backward_demodulation, [], [f160, f504])).
fof(f160, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3995, plain, (~ spl18_36 | ~ spl18_40), inference(avatar_split_clause, [], [f3992, f527, f510])).
fof(f510, plain, (spl18_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_36])])).
fof(f3992, plain, (~ (e3 = op(e1, e3)) | ~ spl18_40), inference(backward_demodulation, [], [f174, f529])).
fof(f529, plain, ((e3 = op(e1, e2)) | ~ spl18_40), inference(avatar_component_clause, [], [f527])).
fof(f3994, plain, (~ spl18_24 | ~ spl18_40), inference(avatar_split_clause, [], [f3990, f527, f459])).
fof(f3990, plain, (~ (e3 = op(e2, e2)) | ~ spl18_40), inference(backward_demodulation, [], [f154, f529])).
fof(f3989, plain, (~ spl18_45 | spl18_89), inference(avatar_contradiction_clause, [], [f3988])).
fof(f3988, plain, ($false | (~ spl18_45 | spl18_89)), inference(subsumption_resolution, [], [f3985, f551])).
fof(f3985, plain, (~ (e0 = op(e1, e0)) | (~ spl18_45 | spl18_89)), inference(backward_demodulation, [], [f866, f551])).
fof(f866, plain, (~ (e0 = op(e1, op(e1, e0))) | spl18_89), inference(avatar_component_clause, [], [f865])).
fof(f865, plain, (spl18_89 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl18_89])])).
fof(f3987, plain, (~ spl18_37 | ~ spl18_45), inference(avatar_split_clause, [], [f3980, f549, f515])).
fof(f515, plain, (spl18_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_37])])).
fof(f3980, plain, (~ (e0 = op(e1, e2)) | ~ spl18_45), inference(backward_demodulation, [], [f170, f551])).
fof(f170, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3986, plain, (~ spl18_13 | ~ spl18_45), inference(avatar_split_clause, [], [f3979, f549, f413])).
fof(f413, plain, (spl18_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_13])])).
fof(f3979, plain, (~ (e0 = op(e3, e0)) | ~ spl18_45), inference(backward_demodulation, [], [f143, f551])).
fof(f143, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3978, plain, (~ spl18_36 | ~ spl18_52), inference(avatar_split_clause, [], [f3975, f578, f510])).
fof(f3975, plain, (~ (e3 = op(e1, e3)) | ~ spl18_52), inference(backward_demodulation, [], [f157, f580])).
fof(f3974, plain, (~ spl18_49 | ~ spl18_57), inference(avatar_split_clause, [], [f3968, f600, f566])).
fof(f566, plain, (spl18_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_49])])).
fof(f3968, plain, (~ (e0 = op(e0, e3)) | ~ spl18_57), inference(backward_demodulation, [], [f167, f602])).
fof(f167, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3973, plain, (~ spl18_25 | ~ spl18_57), inference(avatar_split_clause, [], [f3966, f600, f464])).
fof(f464, plain, (spl18_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_25])])).
fof(f3966, plain, (~ (e0 = op(e2, e1)) | ~ spl18_57), inference(backward_demodulation, [], [f146, f602])).
fof(f146, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3965, plain, (~ spl18_28 | ~ spl18_62 | spl18_113 | ~ spl18_114), inference(avatar_split_clause, [], [f3961, f993, f989, f621, f476])).
fof(f989, plain, (spl18_113 <=> (e3 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl18_113])])).
fof(f3961, plain, (~ (e3 = op(e2, e1)) | (~ spl18_62 | spl18_113 | ~ spl18_114)), inference(backward_demodulation, [], [f3940, f623])).
fof(f3940, plain, (~ (e3 = op(e2, op(e0, e0))) | (spl18_113 | ~ spl18_114)), inference(forward_demodulation, [], [f991, f994])).
fof(f991, plain, (~ (e3 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | spl18_113), inference(avatar_component_clause, [], [f989])).
fof(f3964, plain, (spl18_57 | ~ spl18_62 | ~ spl18_94), inference(avatar_split_clause, [], [f3958, f892, f621, f600])).
fof(f892, plain, (spl18_94 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl18_94])])).
fof(f3958, plain, ((e0 = op(e0, e1)) | (~ spl18_62 | ~ spl18_94)), inference(backward_demodulation, [], [f894, f623])).
fof(f894, plain, ((e0 = op(e0, op(e0, e0))) | ~ spl18_94), inference(avatar_component_clause, [], [f892])).
fof(f3963, plain, (~ spl18_58 | ~ spl18_62), inference(avatar_split_clause, [], [f3954, f621, f604])).
fof(f3954, plain, (~ (e1 = op(e0, e1)) | ~ spl18_62), inference(backward_demodulation, [], [f163, f623])).
fof(f163, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f3962, plain, (~ spl18_46 | ~ spl18_62), inference(avatar_split_clause, [], [f3952, f621, f553])).
fof(f3952, plain, (~ (e1 = op(e1, e0)) | ~ spl18_62), inference(backward_demodulation, [], [f139, f623])).
fof(f139, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f3949, plain, (~ spl18_24 | ~ spl18_43 | spl18_100), inference(avatar_split_clause, [], [f3948, f930, f540, f459])).
fof(f540, plain, (spl18_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_43])])).
fof(f930, plain, (spl18_100 <=> (e3 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl18_100])])).
fof(f3948, plain, (~ (e3 = op(e2, e2)) | (~ spl18_43 | spl18_100)), inference(forward_demodulation, [], [f932, f542])).
fof(f542, plain, ((e2 = op(e1, e1)) | ~ spl18_43), inference(avatar_component_clause, [], [f540])).
fof(f932, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | spl18_100), inference(avatar_component_clause, [], [f930])).
fof(f3925, plain, (~ spl18_5 | ~ spl18_13), inference(avatar_split_clause, [], [f3921, f413, f379])).
fof(f3921, plain, (~ (e0 = op(e3, e2)) | ~ spl18_13), inference(backward_demodulation, [], [f182, f415])).
fof(f415, plain, ((e0 = op(e3, e0)) | ~ spl18_13), inference(avatar_component_clause, [], [f413])).
fof(f3915, plain, (~ spl18_17 | ~ spl18_25), inference(avatar_split_clause, [], [f3910, f464, f430])).
fof(f3910, plain, (~ (e0 = op(e2, e3)) | ~ spl18_25), inference(backward_demodulation, [], [f179, f466])).
fof(f466, plain, ((e0 = op(e2, e1)) | ~ spl18_25), inference(avatar_component_clause, [], [f464])).
fof(f179, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3900, plain, (~ spl18_5 | ~ spl18_37), inference(avatar_split_clause, [], [f3894, f515, f379])).
fof(f3894, plain, (~ (e0 = op(e3, e2)) | ~ spl18_37), inference(backward_demodulation, [], [f155, f517])).
fof(f517, plain, ((e0 = op(e1, e2)) | ~ spl18_37), inference(avatar_component_clause, [], [f515])).
fof(f155, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3893, plain, (~ spl18_5 | ~ spl18_43 | ~ spl18_100 | spl18_121), inference(avatar_split_clause, [], [f3889, f1028, f930, f540, f379])).
fof(f1028, plain, (spl18_121 <=> (e0 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl18_121])])).
fof(f3889, plain, (~ (e0 = op(e3, e2)) | (~ spl18_43 | ~ spl18_100 | spl18_121)), inference(backward_demodulation, [], [f3831, f542])).
fof(f3831, plain, (~ (e0 = op(e3, op(e1, e1))) | (~ spl18_100 | spl18_121)), inference(backward_demodulation, [], [f1030, f931])).
fof(f931, plain, ((e3 = op(op(e1, e1), op(e1, e1))) | ~ spl18_100), inference(avatar_component_clause, [], [f930])).
fof(f1030, plain, (~ (e0 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | spl18_121), inference(avatar_component_clause, [], [f1028])).
fof(f3892, plain, (~ spl18_39 | ~ spl18_43), inference(avatar_split_clause, [], [f3882, f540, f523])).
fof(f3882, plain, (~ (e2 = op(e1, e2)) | ~ spl18_43), inference(backward_demodulation, [], [f172, f542])).
fof(f3891, plain, (~ spl18_27 | ~ spl18_43), inference(avatar_split_clause, [], [f3880, f540, f472])).
fof(f3880, plain, (~ (e2 = op(e2, e1)) | ~ spl18_43), inference(backward_demodulation, [], [f148, f542])).
fof(f3878, plain, (~ spl18_30 | ~ spl18_46), inference(avatar_split_clause, [], [f3870, f553, f485])).
fof(f3870, plain, (~ (e1 = op(e2, e0)) | ~ spl18_46), inference(backward_demodulation, [], [f142, f555])).
fof(f3869, plain, (~ spl18_17 | ~ spl18_49), inference(avatar_split_clause, [], [f3864, f566, f430])).
fof(f3864, plain, (~ (e0 = op(e2, e3)) | ~ spl18_49), inference(backward_demodulation, [], [f158, f568])).
fof(f568, plain, ((e0 = op(e0, e3)) | ~ spl18_49), inference(avatar_component_clause, [], [f566])).
fof(f3856, plain, (~ spl18_54 | ~ spl18_58), inference(avatar_split_clause, [], [f3850, f604, f587])).
fof(f3850, plain, (~ (e1 = op(e0, e2)) | ~ spl18_58), inference(backward_demodulation, [], [f166, f606])).
fof(f3846, plain, (~ spl18_18 | ~ spl18_64 | ~ spl18_114 | spl18_119), inference(avatar_split_clause, [], [f3842, f1018, f993, f629, f434])).
fof(f1018, plain, (spl18_119 <=> (e1 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl18_119])])).
fof(f3842, plain, (~ (e1 = op(e2, e3)) | (~ spl18_64 | ~ spl18_114 | spl18_119)), inference(backward_demodulation, [], [f3827, f631])).
fof(f3827, plain, (~ (e1 = op(e2, op(e0, e0))) | (~ spl18_114 | spl18_119)), inference(backward_demodulation, [], [f1020, f994])).
fof(f1020, plain, (~ (e1 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | spl18_119), inference(avatar_component_clause, [], [f1018])).
fof(f3845, plain, (spl18_49 | ~ spl18_64 | ~ spl18_94), inference(avatar_split_clause, [], [f3839, f892, f629, f566])).
fof(f3839, plain, ((e0 = op(e0, e3)) | (~ spl18_64 | ~ spl18_94)), inference(backward_demodulation, [], [f894, f631])).
fof(f3816, plain, (~ spl18_19 | ~ spl18_3), inference(avatar_split_clause, [], [f3815, f370, f438])).
fof(f438, plain, (spl18_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_19])])).
fof(f3815, plain, (~ (e2 = op(e2, e3)) | ~ spl18_3), inference(forward_demodulation, [], [f162, f372])).
fof(f372, plain, ((e2 = op(e3, e3)) | ~ spl18_3), inference(avatar_component_clause, [], [f370])).
fof(f3810, plain, (~ spl18_7 | ~ spl18_3), inference(avatar_split_clause, [], [f3809, f370, f387])).
fof(f387, plain, (spl18_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_7])])).
fof(f3809, plain, (~ (e2 = op(e3, e2)) | ~ spl18_3), inference(forward_demodulation, [], [f186, f372])).
fof(f3808, plain, (~ spl18_8 | ~ spl18_3 | spl18_95), inference(avatar_split_clause, [], [f3766, f906, f370, f391])).
fof(f906, plain, (spl18_95 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl18_95])])).
fof(f3766, plain, (~ (e3 = op(e3, e2)) | (~ spl18_3 | spl18_95)), inference(backward_demodulation, [], [f907, f372])).
fof(f907, plain, (~ (e3 = op(e3, op(e3, e3))) | spl18_95), inference(avatar_component_clause, [], [f906])).
fof(f3779, plain, (~ spl18_3 | ~ spl18_17 | ~ spl18_44 | spl18_121), inference(avatar_contradiction_clause, [], [f3778])).
fof(f3778, plain, ($false | (~ spl18_3 | ~ spl18_17 | ~ spl18_44 | spl18_121)), inference(subsumption_resolution, [], [f3771, f432])).
fof(f3771, plain, (~ (e0 = op(e2, e3)) | (~ spl18_3 | ~ spl18_44 | spl18_121)), inference(backward_demodulation, [], [f3721, f372])).
fof(f3721, plain, (~ (e0 = op(op(e3, e3), e3)) | (~ spl18_44 | spl18_121)), inference(backward_demodulation, [], [f1030, f546])).
fof(f3775, plain, (~ spl18_3 | spl18_23 | ~ spl18_108), inference(avatar_contradiction_clause, [], [f3774])).
fof(f3774, plain, ($false | (~ spl18_3 | spl18_23 | ~ spl18_108)), inference(subsumption_resolution, [], [f3767, f456])).
fof(f456, plain, (~ (e2 = op(e2, e2)) | spl18_23), inference(avatar_component_clause, [], [f455])).
fof(f3767, plain, ((e2 = op(e2, e2)) | (~ spl18_3 | ~ spl18_108)), inference(backward_demodulation, [], [f967, f372])).
fof(f967, plain, ((e2 = op(op(e3, e3), op(e3, e3))) | ~ spl18_108), inference(avatar_component_clause, [], [f966])).
fof(f966, plain, (spl18_108 <=> (e2 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl18_108])])).
fof(f3763, plain, (~ spl18_5 | spl18_15 | ~ spl18_96), inference(avatar_contradiction_clause, [], [f3762])).
fof(f3762, plain, ($false | (~ spl18_5 | spl18_15 | ~ spl18_96)), inference(subsumption_resolution, [], [f3759, f422])).
fof(f422, plain, (~ (e2 = op(e3, e0)) | spl18_15), inference(avatar_component_clause, [], [f421])).
fof(f3759, plain, ((e2 = op(e3, e0)) | (~ spl18_5 | ~ spl18_96)), inference(backward_demodulation, [], [f913, f381])).
fof(f913, plain, ((e2 = op(e3, op(e3, e2))) | ~ spl18_96), inference(avatar_component_clause, [], [f911])).
fof(f911, plain, (spl18_96 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl18_96])])).
fof(f3745, plain, (~ spl18_19 | ~ spl18_24 | spl18_82), inference(avatar_split_clause, [], [f3738, f828, f459, f438])).
fof(f3738, plain, (~ (e2 = op(e2, e3)) | (~ spl18_24 | spl18_82)), inference(backward_demodulation, [], [f829, f461])).
fof(f829, plain, (~ (e2 = op(e2, op(e2, e2))) | spl18_82), inference(avatar_component_clause, [], [f828])).
fof(f3735, plain, (~ spl18_19 | ~ spl18_27), inference(avatar_split_clause, [], [f3732, f472, f438])).
fof(f3732, plain, (~ (e2 = op(e2, e3)) | ~ spl18_27), inference(backward_demodulation, [], [f179, f474])).
fof(f3730, plain, (~ spl18_7 | ~ spl18_39), inference(avatar_split_clause, [], [f3727, f523, f387])).
fof(f3727, plain, (~ (e2 = op(e3, e2)) | ~ spl18_39), inference(backward_demodulation, [], [f155, f525])).
fof(f3714, plain, (spl18_44 | ~ spl18_34 | ~ spl18_86), inference(avatar_split_clause, [], [f3713, f850, f502, f544])).
fof(f850, plain, (spl18_86 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl18_86])])).
fof(f3713, plain, ((e3 = op(e1, e1)) | (~ spl18_34 | ~ spl18_86)), inference(forward_demodulation, [], [f852, f504])).
fof(f852, plain, ((e3 = op(e1, op(e1, e3))) | ~ spl18_86), inference(avatar_component_clause, [], [f850])).
fof(f3709, plain, (spl18_24 | ~ spl18_63 | ~ spl18_112), inference(avatar_split_clause, [], [f3708, f984, f625, f459])).
fof(f625, plain, (spl18_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl18_63])])).
fof(f3708, plain, ((e3 = op(e2, e2)) | (~ spl18_63 | ~ spl18_112)), inference(forward_demodulation, [], [f985, f627])).
fof(f627, plain, ((op(e0, e0) = e2) | ~ spl18_63), inference(avatar_component_clause, [], [f625])).
fof(f3702, plain, (~ spl18_1 | ~ spl18_30 | ~ spl18_63 | spl18_107), inference(avatar_contradiction_clause, [], [f3701])).
fof(f3701, plain, ($false | (~ spl18_1 | ~ spl18_30 | ~ spl18_63 | spl18_107)), inference(subsumption_resolution, [], [f3700, f487])).
fof(f3700, plain, (~ (e1 = op(e2, e0)) | (~ spl18_1 | ~ spl18_63 | spl18_107)), inference(forward_demodulation, [], [f3699, f627])).
fof(f3699, plain, (~ (e1 = op(op(e0, e0), e0)) | (~ spl18_1 | spl18_107)), inference(forward_demodulation, [], [f964, f364])).
fof(f964, plain, (~ (e1 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | spl18_107), inference(avatar_component_clause, [], [f962])).
fof(f962, plain, (spl18_107 <=> (e1 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl18_107])])).
fof(f3687, plain, (~ spl18_17 | ~ spl18_21), inference(avatar_split_clause, [], [f3675, f447, f430])).
fof(f3675, plain, (~ (e0 = op(e2, e3)) | ~ spl18_21), inference(backward_demodulation, [], [f180, f449])).
fof(f3686, plain, (~ spl18_5 | ~ spl18_21), inference(avatar_split_clause, [], [f3674, f447, f379])).
fof(f3674, plain, (~ (e0 = op(e3, e2)) | ~ spl18_21), inference(backward_demodulation, [], [f156, f449])).
fof(f3668, plain, (~ spl18_23 | ~ spl18_63 | spl18_114), inference(avatar_split_clause, [], [f3667, f993, f625, f455])).
fof(f3667, plain, (~ (e2 = op(e2, e2)) | (~ spl18_63 | spl18_114)), inference(forward_demodulation, [], [f995, f627])).
fof(f995, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | spl18_114), inference(avatar_component_clause, [], [f993])).
fof(f3654, plain, (spl18_16 | ~ spl18_1 | ~ spl18_95), inference(avatar_split_clause, [], [f3653, f906, f362, f425])).
fof(f3653, plain, ((e3 = op(e3, e0)) | (~ spl18_1 | ~ spl18_95)), inference(forward_demodulation, [], [f908, f364])).
fof(f908, plain, ((e3 = op(e3, op(e3, e3))) | ~ spl18_95), inference(avatar_component_clause, [], [f906])).
fof(f3652, plain, (~ spl18_1 | ~ spl18_63 | spl18_108), inference(avatar_contradiction_clause, [], [f3651])).
fof(f3651, plain, ($false | (~ spl18_1 | ~ spl18_63 | spl18_108)), inference(subsumption_resolution, [], [f3650, f627])).
fof(f3650, plain, (~ (op(e0, e0) = e2) | (~ spl18_1 | spl18_108)), inference(forward_demodulation, [], [f968, f364])).
fof(f968, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | spl18_108), inference(avatar_component_clause, [], [f966])).
fof(f3645, plain, (~ spl18_1 | ~ spl18_4), inference(avatar_contradiction_clause, [], [f3644])).
fof(f3644, plain, ($false | (~ spl18_1 | ~ spl18_4)), inference(subsumption_resolution, [], [f3643, f189])).
fof(f3643, plain, ((e0 = e3) | (~ spl18_1 | ~ spl18_4)), inference(backward_demodulation, [], [f376, f364])).
fof(f376, plain, ((e3 = op(e3, e3)) | ~ spl18_4), inference(avatar_component_clause, [], [f374])).
fof(f374, plain, (spl18_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl18_4])])).
fof(f3621, plain, (~ spl18_15 | ~ spl18_16), inference(avatar_contradiction_clause, [], [f3620])).
fof(f3620, plain, ($false | (~ spl18_15 | ~ spl18_16)), inference(subsumption_resolution, [], [f3619, f192])).
fof(f3619, plain, ((e2 = e3) | (~ spl18_15 | ~ spl18_16)), inference(backward_demodulation, [], [f427, f423])).
fof(f3618, plain, (spl18_1 | ~ spl18_16 | ~ spl18_98), inference(avatar_split_clause, [], [f3615, f921, f425, f362])).
fof(f921, plain, (spl18_98 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl18_98])])).
fof(f3615, plain, ((e0 = op(e3, e3)) | (~ spl18_16 | ~ spl18_98)), inference(backward_demodulation, [], [f923, f427])).
fof(f923, plain, ((e0 = op(e3, op(e3, e0))) | ~ spl18_98), inference(avatar_component_clause, [], [f921])).
fof(f3617, plain, (~ spl18_4 | ~ spl18_16), inference(avatar_split_clause, [], [f3614, f425, f374])).
fof(f3614, plain, (~ (e3 = op(e3, e3)) | ~ spl18_16), inference(backward_demodulation, [], [f183, f427])).
fof(f3601, plain, (~ spl18_23 | spl18_82), inference(avatar_contradiction_clause, [], [f3600])).
fof(f3600, plain, ($false | (~ spl18_23 | spl18_82)), inference(subsumption_resolution, [], [f3595, f457])).
fof(f457, plain, ((e2 = op(e2, e2)) | ~ spl18_23), inference(avatar_component_clause, [], [f455])).
fof(f3595, plain, (~ (e2 = op(e2, e2)) | (~ spl18_23 | spl18_82)), inference(backward_demodulation, [], [f829, f457])).
fof(f3591, plain, (~ spl18_20 | ~ spl18_28), inference(avatar_split_clause, [], [f3589, f476, f442])).
fof(f3589, plain, (~ (e3 = op(e2, e3)) | ~ spl18_28), inference(backward_demodulation, [], [f179, f478])).
fof(f3586, plain, (~ spl18_8 | ~ spl18_40), inference(avatar_split_clause, [], [f3584, f527, f391])).
fof(f3584, plain, (~ (e3 = op(e3, e2)) | ~ spl18_40), inference(backward_demodulation, [], [f155, f529])).
fof(f3582, plain, (~ spl18_21 | ~ spl18_43 | spl18_122), inference(avatar_split_clause, [], [f3577, f1033, f540, f447])).
fof(f3577, plain, (~ (e0 = op(e2, e2)) | (~ spl18_43 | spl18_122)), inference(backward_demodulation, [], [f1035, f542])).
fof(f3581, plain, (spl18_23 | ~ spl18_43 | ~ spl18_102), inference(avatar_split_clause, [], [f3576, f939, f540, f455])).
fof(f3576, plain, ((e2 = op(e2, e2)) | (~ spl18_43 | ~ spl18_102)), inference(backward_demodulation, [], [f940, f542])).
fof(f940, plain, ((e2 = op(op(e1, e1), op(e1, e1))) | ~ spl18_102), inference(avatar_component_clause, [], [f939])).
fof(f3571, plain, (~ spl18_41 | ~ spl18_45), inference(avatar_split_clause, [], [f3565, f549, f532])).
fof(f532, plain, (spl18_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_41])])).
fof(f3565, plain, (~ (e0 = op(e1, e1)) | ~ spl18_45), inference(backward_demodulation, [], [f169, f551])).
fof(f169, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3563, plain, (~ spl18_51 | ~ spl18_52), inference(avatar_contradiction_clause, [], [f3562])).
fof(f3562, plain, ($false | (~ spl18_51 | ~ spl18_52)), inference(subsumption_resolution, [], [f3561, f192])).
fof(f3561, plain, ((e2 = e3) | (~ spl18_51 | ~ spl18_52)), inference(backward_demodulation, [], [f580, f576])).
fof(f3560, plain, (~ spl18_4 | ~ spl18_52), inference(avatar_split_clause, [], [f3558, f578, f374])).
fof(f3558, plain, (~ (e3 = op(e3, e3)) | ~ spl18_52), inference(backward_demodulation, [], [f159, f580])).
fof(f3554, plain, (~ spl18_41 | ~ spl18_57), inference(avatar_split_clause, [], [f3547, f600, f532])).
fof(f3547, plain, (~ (e0 = op(e1, e1)) | ~ spl18_57), inference(backward_demodulation, [], [f145, f602])).
fof(f145, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3545, plain, (spl18_23 | ~ spl18_63 | ~ spl18_114), inference(avatar_split_clause, [], [f3544, f993, f625, f455])).
fof(f3544, plain, ((e2 = op(e2, e2)) | (~ spl18_63 | ~ spl18_114)), inference(forward_demodulation, [], [f994, f627])).
fof(f3542, plain, (~ spl18_51 | ~ spl18_63), inference(avatar_split_clause, [], [f3417, f625, f574])).
fof(f3417, plain, (~ (e2 = op(e0, e3)) | ~ spl18_63), inference(backward_demodulation, [], [f165, f627])).
fof(f3540, plain, (~ spl18_42 | ~ spl18_10), inference(avatar_split_clause, [], [f3539, f400, f536])).
fof(f536, plain, (spl18_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_42])])).
fof(f3539, plain, (~ (e1 = op(e1, e1)) | ~ spl18_10), inference(forward_demodulation, [], [f149, f402])).
fof(f149, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3538, plain, (~ spl18_42 | ~ spl18_34), inference(avatar_split_clause, [], [f3537, f502, f536])).
fof(f3537, plain, (~ (e1 = op(e1, e1)) | ~ spl18_34), inference(forward_demodulation, [], [f173, f504])).
fof(f173, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3536, plain, (~ spl18_44 | ~ spl18_34 | spl18_86), inference(avatar_split_clause, [], [f3329, f850, f502, f544])).
fof(f3329, plain, (~ (e3 = op(e1, e1)) | (~ spl18_34 | spl18_86)), inference(forward_demodulation, [], [f851, f504])).
fof(f851, plain, (~ (e3 = op(e1, op(e1, e3))) | spl18_86), inference(avatar_component_clause, [], [f850])).
fof(f3534, plain, (~ spl18_15 | ~ spl18_63), inference(avatar_split_clause, [], [f3414, f625, f421])).
fof(f3414, plain, (~ (e2 = op(e3, e0)) | ~ spl18_63), inference(backward_demodulation, [], [f141, f627])).
fof(f3528, plain, (~ spl18_3 | ~ spl18_21 | spl18_118), inference(avatar_contradiction_clause, [], [f3527])).
fof(f3527, plain, ($false | (~ spl18_3 | ~ spl18_21 | spl18_118)), inference(subsumption_resolution, [], [f3526, f449])).
fof(f3526, plain, (~ (e0 = op(e2, e2)) | (~ spl18_3 | spl18_118)), inference(forward_demodulation, [], [f1015, f372])).
fof(f1015, plain, (~ (e0 = op(op(e3, e3), op(e3, e3))) | spl18_118), inference(avatar_component_clause, [], [f1013])).
fof(f1013, plain, (spl18_118 <=> (e0 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl18_118])])).
fof(f3525, plain, (~ spl18_21 | ~ spl18_30 | ~ spl18_63 | spl18_103), inference(avatar_contradiction_clause, [], [f3524])).
fof(f3524, plain, ($false | (~ spl18_21 | ~ spl18_30 | ~ spl18_63 | spl18_103)), inference(subsumption_resolution, [], [f3523, f487])).
fof(f3523, plain, (~ (e1 = op(e2, e0)) | (~ spl18_21 | ~ spl18_63 | spl18_103)), inference(forward_demodulation, [], [f3522, f627])).
fof(f3522, plain, (~ (e1 = op(op(e0, e0), e0)) | (~ spl18_21 | spl18_103)), inference(forward_demodulation, [], [f946, f449])).
fof(f3510, plain, (~ spl18_3 | ~ spl18_21 | ~ spl18_54 | spl18_107), inference(avatar_contradiction_clause, [], [f3509])).
fof(f3509, plain, ($false | (~ spl18_3 | ~ spl18_21 | ~ spl18_54 | spl18_107)), inference(subsumption_resolution, [], [f3508, f589])).
fof(f3508, plain, (~ (e1 = op(e0, e2)) | (~ spl18_3 | ~ spl18_21 | spl18_107)), inference(forward_demodulation, [], [f3505, f449])).
fof(f3505, plain, (~ (e1 = op(op(e2, e2), e2)) | (~ spl18_3 | spl18_107)), inference(backward_demodulation, [], [f964, f372])).
fof(f3502, plain, (spl18_3 | ~ spl18_8 | ~ spl18_96), inference(avatar_split_clause, [], [f3501, f911, f391, f370])).
fof(f3501, plain, ((e2 = op(e3, e3)) | (~ spl18_8 | ~ spl18_96)), inference(backward_demodulation, [], [f913, f393])).
fof(f393, plain, ((e3 = op(e3, e2)) | ~ spl18_8), inference(avatar_component_clause, [], [f391])).
fof(f3499, plain, (~ spl18_1 | ~ spl18_13), inference(avatar_split_clause, [], [f3494, f413, f362])).
fof(f3494, plain, (~ (e0 = op(e3, e3)) | ~ spl18_13), inference(backward_demodulation, [], [f183, f415])).
fof(f3455, plain, (~ spl18_40 | ~ spl18_48), inference(avatar_split_clause, [], [f3452, f561, f527])).
fof(f561, plain, (spl18_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_48])])).
fof(f3452, plain, (~ (e3 = op(e1, e2)) | ~ spl18_48), inference(backward_demodulation, [], [f170, f563])).
fof(f563, plain, ((e3 = op(e1, e0)) | ~ spl18_48), inference(avatar_component_clause, [], [f561])).
fof(f3454, plain, (~ spl18_16 | ~ spl18_48), inference(avatar_split_clause, [], [f3451, f561, f425])).
fof(f3451, plain, (~ (e3 = op(e3, e0)) | ~ spl18_48), inference(backward_demodulation, [], [f143, f563])).
fof(f3448, plain, (~ spl18_1 | ~ spl18_49), inference(avatar_split_clause, [], [f3443, f566, f362])).
fof(f3443, plain, (~ (e0 = op(e3, e3)) | ~ spl18_49), inference(backward_demodulation, [], [f159, f568])).
fof(f3434, plain, (~ spl18_52 | ~ spl18_60), inference(avatar_split_clause, [], [f3431, f612, f578])).
fof(f612, plain, (spl18_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_60])])).
fof(f3431, plain, (~ (e3 = op(e0, e3)) | ~ spl18_60), inference(backward_demodulation, [], [f167, f614])).
fof(f614, plain, ((e3 = op(e0, e1)) | ~ spl18_60), inference(avatar_component_clause, [], [f612])).
fof(f3433, plain, (~ spl18_28 | ~ spl18_60), inference(avatar_split_clause, [], [f3429, f612, f476])).
fof(f3429, plain, (~ (e3 = op(e2, e1)) | ~ spl18_60), inference(backward_demodulation, [], [f146, f614])).
fof(f3428, plain, (~ spl18_22 | ~ spl18_63 | spl18_120), inference(avatar_split_clause, [], [f3422, f1023, f625, f451])).
fof(f3422, plain, (~ (e1 = op(e2, e2)) | (~ spl18_63 | spl18_120)), inference(backward_demodulation, [], [f1025, f627])).
fof(f3426, plain, (~ spl18_59 | ~ spl18_63), inference(avatar_split_clause, [], [f3415, f625, f608])).
fof(f608, plain, (spl18_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl18_59])])).
fof(f3415, plain, (~ (e2 = op(e0, e1)) | ~ spl18_63), inference(backward_demodulation, [], [f163, f627])).
fof(f3425, plain, (~ spl18_47 | ~ spl18_63), inference(avatar_split_clause, [], [f3412, f625, f557])).
fof(f557, plain, (spl18_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_47])])).
fof(f3412, plain, (~ (e2 = op(e1, e0)) | ~ spl18_63), inference(backward_demodulation, [], [f139, f627])).
fof(f3390, plain, (~ spl18_39 | ~ spl18_47), inference(avatar_split_clause, [], [f3386, f557, f523])).
fof(f3386, plain, (~ (e2 = op(e1, e2)) | ~ spl18_47), inference(backward_demodulation, [], [f170, f559])).
fof(f559, plain, ((e2 = op(e1, e0)) | ~ spl18_47), inference(avatar_component_clause, [], [f557])).
fof(f3389, plain, (~ spl18_15 | ~ spl18_47), inference(avatar_split_clause, [], [f3385, f557, f421])).
fof(f3385, plain, (~ (e2 = op(e3, e0)) | ~ spl18_47), inference(backward_demodulation, [], [f143, f559])).
fof(f3379, plain, (~ spl18_27 | ~ spl18_59), inference(avatar_split_clause, [], [f3375, f608, f472])).
fof(f3375, plain, (~ (e2 = op(e2, e1)) | ~ spl18_59), inference(backward_demodulation, [], [f146, f610])).
fof(f610, plain, ((e2 = op(e0, e1)) | ~ spl18_59), inference(avatar_component_clause, [], [f608])).
fof(f3374, plain, (~ spl18_27 | ~ spl18_22 | spl18_82), inference(avatar_split_clause, [], [f3373, f828, f451, f472])).
fof(f3373, plain, (~ (e2 = op(e2, e1)) | (~ spl18_22 | spl18_82)), inference(forward_demodulation, [], [f829, f453])).
fof(f3363, plain, (~ spl18_60 | ~ spl18_22 | spl18_105 | ~ spl18_116), inference(avatar_split_clause, [], [f3337, f1003, f953, f451, f612])).
fof(f953, plain, (spl18_105 <=> (e3 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl18_105])])).
fof(f3337, plain, (~ (e3 = op(e0, e1)) | (~ spl18_22 | spl18_105 | ~ spl18_116)), inference(backward_demodulation, [], [f3322, f453])).
fof(f3322, plain, (~ (e3 = op(e0, op(e2, e2))) | (spl18_105 | ~ spl18_116)), inference(backward_demodulation, [], [f955, f1004])).
fof(f1004, plain, ((e0 = op(op(e2, e2), op(e2, e2))) | ~ spl18_116), inference(avatar_component_clause, [], [f1003])).
fof(f955, plain, (~ (e3 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | spl18_105), inference(avatar_component_clause, [], [f953])).
fof(f3362, plain, (~ spl18_59 | ~ spl18_41 | ~ spl18_62 | spl18_111), inference(avatar_split_clause, [], [f3200, f980, f621, f532, f608])).
fof(f3200, plain, (~ (e2 = op(e0, e1)) | (~ spl18_41 | ~ spl18_62 | spl18_111)), inference(forward_demodulation, [], [f3190, f534])).
fof(f534, plain, ((e0 = op(e1, e1)) | ~ spl18_41), inference(avatar_component_clause, [], [f532])).
fof(f3190, plain, (~ (e2 = op(op(e1, e1), e1)) | (~ spl18_62 | spl18_111)), inference(backward_demodulation, [], [f982, f623])).
fof(f3354, plain, (~ spl18_47 | ~ spl18_1 | ~ spl18_62 | spl18_109), inference(avatar_split_clause, [], [f3315, f971, f621, f362, f557])).
fof(f971, plain, (spl18_109 <=> (e2 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl18_109])])).
fof(f3315, plain, (~ (e2 = op(e1, e0)) | (~ spl18_1 | ~ spl18_62 | spl18_109)), inference(forward_demodulation, [], [f3314, f623])).
fof(f3314, plain, (~ (e2 = op(op(e0, e0), e0)) | (~ spl18_1 | spl18_109)), inference(forward_demodulation, [], [f973, f364])).
fof(f973, plain, (~ (e2 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | spl18_109), inference(avatar_component_clause, [], [f971])).
fof(f3318, plain, (~ spl18_1 | ~ spl18_62 | spl18_110), inference(avatar_contradiction_clause, [], [f3317])).
fof(f3317, plain, ($false | (~ spl18_1 | ~ spl18_62 | spl18_110)), inference(subsumption_resolution, [], [f3316, f623])).
fof(f3316, plain, (~ (op(e0, e0) = e1) | (~ spl18_1 | spl18_110)), inference(forward_demodulation, [], [f977, f364])).
fof(f977, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | spl18_110), inference(avatar_component_clause, [], [f975])).
fof(f975, plain, (spl18_110 <=> (e1 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl18_110])])).
fof(f3301, plain, (spl18_22 | ~ spl18_27 | ~ spl18_83), inference(avatar_split_clause, [], [f3174, f833, f472, f451])).
fof(f833, plain, (spl18_83 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl18_83])])).
fof(f3174, plain, ((e1 = op(e2, e2)) | (~ spl18_27 | ~ spl18_83)), inference(forward_demodulation, [], [f835, f474])).
fof(f835, plain, ((e1 = op(e2, op(e2, e1))) | ~ spl18_83), inference(avatar_component_clause, [], [f833])).
fof(f3287, plain, (~ spl18_1 | spl18_61 | ~ spl18_118), inference(avatar_contradiction_clause, [], [f3286])).
fof(f3286, plain, ($false | (~ spl18_1 | spl18_61 | ~ spl18_118)), inference(subsumption_resolution, [], [f3283, f618])).
fof(f618, plain, (~ (e0 = op(e0, e0)) | spl18_61), inference(avatar_component_clause, [], [f617])).
fof(f617, plain, (spl18_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_61])])).
fof(f3283, plain, ((e0 = op(e0, e0)) | (~ spl18_1 | ~ spl18_118)), inference(backward_demodulation, [], [f1014, f364])).
fof(f1014, plain, ((e0 = op(op(e3, e3), op(e3, e3))) | ~ spl18_118), inference(avatar_component_clause, [], [f1013])).
fof(f3271, plain, (spl18_5 | ~ spl18_15 | ~ spl18_98), inference(avatar_contradiction_clause, [], [f3270])).
fof(f3270, plain, ($false | (spl18_5 | ~ spl18_15 | ~ spl18_98)), inference(subsumption_resolution, [], [f3268, f380])).
fof(f380, plain, (~ (e0 = op(e3, e2)) | spl18_5), inference(avatar_component_clause, [], [f379])).
fof(f3268, plain, ((e0 = op(e3, e2)) | (~ spl18_15 | ~ spl18_98)), inference(backward_demodulation, [], [f923, f423])).
fof(f3263, plain, (~ spl18_18 | ~ spl18_20), inference(avatar_contradiction_clause, [], [f3262])).
fof(f3262, plain, ($false | (~ spl18_18 | ~ spl18_20)), inference(subsumption_resolution, [], [f3261, f191])).
fof(f3261, plain, ((e1 = e3) | (~ spl18_18 | ~ spl18_20)), inference(forward_demodulation, [], [f444, f436])).
fof(f436, plain, ((e1 = op(e2, e3)) | ~ spl18_18), inference(avatar_component_clause, [], [f434])).
fof(f444, plain, ((e3 = op(e2, e3)) | ~ spl18_20), inference(avatar_component_clause, [], [f442])).
fof(f3249, plain, (~ spl18_21 | ~ spl18_29), inference(avatar_split_clause, [], [f3243, f481, f447])).
fof(f3243, plain, (~ (e0 = op(e2, e2)) | ~ spl18_29), inference(backward_demodulation, [], [f176, f483])).
fof(f3248, plain, (~ spl18_13 | ~ spl18_29), inference(avatar_split_clause, [], [f3242, f481, f413])).
fof(f3242, plain, (~ (e0 = op(e3, e0)) | ~ spl18_29), inference(backward_demodulation, [], [f144, f483])).
fof(f3241, plain, (~ spl18_34 | ~ spl18_41 | ~ spl18_86), inference(avatar_contradiction_clause, [], [f3240])).
fof(f3240, plain, ($false | (~ spl18_34 | ~ spl18_41 | ~ spl18_86)), inference(subsumption_resolution, [], [f3239, f189])).
fof(f3239, plain, ((e0 = e3) | (~ spl18_34 | ~ spl18_41 | ~ spl18_86)), inference(forward_demodulation, [], [f3238, f534])).
fof(f3238, plain, ((e3 = op(e1, e1)) | (~ spl18_34 | ~ spl18_86)), inference(backward_demodulation, [], [f852, f504])).
fof(f3234, plain, (~ spl18_36 | ~ spl18_48), inference(avatar_split_clause, [], [f3231, f561, f510])).
fof(f3231, plain, (~ (e3 = op(e1, e3)) | ~ spl18_48), inference(backward_demodulation, [], [f171, f563])).
fof(f3233, plain, (~ spl18_32 | ~ spl18_48), inference(avatar_split_clause, [], [f3229, f561, f493])).
fof(f493, plain, (spl18_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl18_32])])).
fof(f3229, plain, (~ (e3 = op(e2, e0)) | ~ spl18_48), inference(backward_demodulation, [], [f142, f563])).
fof(f3222, plain, (~ spl18_53 | ~ spl18_54), inference(avatar_contradiction_clause, [], [f3221])).
fof(f3221, plain, ($false | (~ spl18_53 | ~ spl18_54)), inference(subsumption_resolution, [], [f3220, f187])).
fof(f187, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f3220, plain, ((e0 = e1) | (~ spl18_53 | ~ spl18_54)), inference(backward_demodulation, [], [f589, f585])).
fof(f3219, plain, (~ spl18_6 | ~ spl18_54), inference(avatar_split_clause, [], [f3214, f587, f383])).
fof(f3214, plain, (~ (e1 = op(e3, e2)) | ~ spl18_54), inference(backward_demodulation, [], [f153, f589])).
fof(f3212, plain, (~ spl18_56 | ~ spl18_60), inference(avatar_split_clause, [], [f3208, f612, f595])).
fof(f595, plain, (spl18_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl18_56])])).
fof(f3208, plain, (~ (e3 = op(e0, e2)) | ~ spl18_60), inference(backward_demodulation, [], [f166, f614])).
fof(f3211, plain, (~ spl18_12 | ~ spl18_60), inference(avatar_split_clause, [], [f3207, f612, f408])).
fof(f3207, plain, (~ (e3 = op(e3, e1)) | ~ spl18_60), inference(backward_demodulation, [], [f147, f614])).
fof(f3199, plain, (~ spl18_54 | ~ spl18_62), inference(avatar_split_clause, [], [f3185, f621, f587])).
fof(f3185, plain, (~ (e1 = op(e0, e2)) | ~ spl18_62), inference(backward_demodulation, [], [f164, f623])).
fof(f164, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3173, plain, (~ spl18_46 | ~ spl18_41 | spl18_88), inference(avatar_split_clause, [], [f3172, f860, f532, f553])).
fof(f3172, plain, (~ (e1 = op(e1, e0)) | (~ spl18_41 | spl18_88)), inference(forward_demodulation, [], [f861, f534])).
fof(f3170, plain, (~ spl18_64 | ~ spl18_41 | spl18_100), inference(avatar_split_clause, [], [f3098, f930, f532, f629])).
fof(f3098, plain, (~ (op(e0, e0) = e3) | (~ spl18_41 | spl18_100)), inference(backward_demodulation, [], [f932, f534])).
fof(f3166, plain, (~ spl18_50 | ~ spl18_18), inference(avatar_split_clause, [], [f3165, f434, f570])).
fof(f3165, plain, (~ (e1 = op(e0, e3)) | ~ spl18_18), inference(forward_demodulation, [], [f158, f436])).
fof(f3145, plain, (~ spl18_2 | ~ spl18_6), inference(avatar_split_clause, [], [f3144, f383, f366])).
fof(f3144, plain, (~ (e1 = op(e3, e3)) | ~ spl18_6), inference(backward_demodulation, [], [f186, f385])).
fof(f385, plain, ((e1 = op(e3, e2)) | ~ spl18_6), inference(avatar_component_clause, [], [f383])).
fof(f3140, plain, (~ spl18_2 | ~ spl18_18), inference(avatar_split_clause, [], [f3138, f434, f366])).
fof(f3138, plain, (~ (e1 = op(e3, e3)) | ~ spl18_18), inference(backward_demodulation, [], [f162, f436])).
fof(f3108, plain, (~ spl18_41 | ~ spl18_63 | spl18_102), inference(avatar_contradiction_clause, [], [f3107])).
fof(f3107, plain, ($false | (~ spl18_41 | ~ spl18_63 | spl18_102)), inference(subsumption_resolution, [], [f3100, f627])).
fof(f3100, plain, (~ (op(e0, e0) = e2) | (~ spl18_41 | spl18_102)), inference(backward_demodulation, [], [f941, f534])).
fof(f3103, plain, (~ spl18_37 | ~ spl18_41), inference(avatar_split_clause, [], [f3092, f532, f515])).
fof(f3092, plain, (~ (e0 = op(e1, e2)) | ~ spl18_41), inference(backward_demodulation, [], [f172, f534])).
fof(f3102, plain, (~ spl18_25 | ~ spl18_41), inference(avatar_split_clause, [], [f3091, f532, f464])).
fof(f3091, plain, (~ (e0 = op(e2, e1)) | ~ spl18_41), inference(backward_demodulation, [], [f148, f534])).
fof(f3089, plain, (spl18_41 | ~ spl18_46 | ~ spl18_89), inference(avatar_split_clause, [], [f3088, f865, f553, f532])).
fof(f3088, plain, ((e0 = op(e1, e1)) | (~ spl18_46 | ~ spl18_89)), inference(forward_demodulation, [], [f867, f555])).
fof(f867, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl18_89), inference(avatar_component_clause, [], [f865])).
fof(f3079, plain, (~ spl18_4 | ~ spl18_12), inference(avatar_split_clause, [], [f3078, f408, f374])).
fof(f3078, plain, (~ (e3 = op(e3, e3)) | ~ spl18_12), inference(forward_demodulation, [], [f185, f410])).
fof(f410, plain, ((e3 = op(e3, e1)) | ~ spl18_12), inference(avatar_component_clause, [], [f408])).
fof(f3046, plain, (~ spl18_21 | ~ spl18_37), inference(avatar_split_clause, [], [f3043, f515, f447])).
fof(f3043, plain, (~ (e0 = op(e2, e2)) | ~ spl18_37), inference(backward_demodulation, [], [f154, f517])).
fof(f3038, plain, (~ spl18_25 | ~ spl18_2 | ~ spl18_43 | spl18_117), inference(avatar_split_clause, [], [f3033, f1008, f540, f366, f464])).
fof(f1008, plain, (spl18_117 <=> (e0 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl18_117])])).
fof(f3033, plain, (~ (e0 = op(e2, e1)) | (~ spl18_2 | ~ spl18_43 | spl18_117)), inference(backward_demodulation, [], [f2831, f542])).
fof(f2831, plain, (~ (e0 = op(op(e1, e1), e1)) | (~ spl18_2 | spl18_117)), inference(backward_demodulation, [], [f1010, f368])).
fof(f368, plain, ((e1 = op(e3, e3)) | ~ spl18_2), inference(avatar_component_clause, [], [f366])).
fof(f1010, plain, (~ (e0 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | spl18_117), inference(avatar_component_clause, [], [f1008])).
fof(f3013, plain, (~ spl18_42 | ~ spl18_46), inference(avatar_split_clause, [], [f3012, f553, f536])).
fof(f3012, plain, (~ (e1 = op(e1, e1)) | ~ spl18_46), inference(forward_demodulation, [], [f169, f555])).
fof(f3011, plain, (~ spl18_42 | ~ spl18_58), inference(avatar_split_clause, [], [f3010, f604, f536])).
fof(f3010, plain, (~ (e1 = op(e1, e1)) | ~ spl18_58), inference(forward_demodulation, [], [f145, f606])).
fof(f3008, plain, (~ spl18_41 | ~ spl18_46 | spl18_89), inference(avatar_split_clause, [], [f2980, f865, f553, f532])).
fof(f2980, plain, (~ (e0 = op(e1, e1)) | (~ spl18_46 | spl18_89)), inference(backward_demodulation, [], [f866, f555])).
fof(f3007, plain, (~ spl18_24 | ~ spl18_63 | spl18_112), inference(avatar_split_clause, [], [f2956, f984, f625, f459])).
fof(f2956, plain, (~ (e3 = op(e2, e2)) | (~ spl18_63 | spl18_112)), inference(backward_demodulation, [], [f986, f627])).
fof(f986, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | spl18_112), inference(avatar_component_clause, [], [f984])).
fof(f3006, plain, (~ spl18_6 | spl18_11 | ~ spl18_96), inference(avatar_contradiction_clause, [], [f3005])).
fof(f3005, plain, ($false | (~ spl18_6 | spl18_11 | ~ spl18_96)), inference(subsumption_resolution, [], [f3004, f405])).
fof(f405, plain, (~ (e2 = op(e3, e1)) | spl18_11), inference(avatar_component_clause, [], [f404])).
fof(f3004, plain, ((e2 = op(e3, e1)) | (~ spl18_6 | ~ spl18_96)), inference(backward_demodulation, [], [f913, f385])).
fof(f3001, plain, (~ spl18_13 | ~ spl18_15), inference(avatar_contradiction_clause, [], [f3000])).
fof(f3000, plain, ($false | (~ spl18_13 | ~ spl18_15)), inference(subsumption_resolution, [], [f2999, f188])).
fof(f188, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f2999, plain, ((e0 = e2) | (~ spl18_13 | ~ spl18_15)), inference(forward_demodulation, [], [f423, f415])).
fof(f2981, plain, (~ spl18_38 | ~ spl18_46), inference(avatar_split_clause, [], [f2977, f553, f519])).
fof(f2977, plain, (~ (e1 = op(e1, e2)) | ~ spl18_46), inference(backward_demodulation, [], [f170, f555])).
fof(f2976, plain, (~ spl18_49 | ~ spl18_51), inference(avatar_contradiction_clause, [], [f2975])).
fof(f2975, plain, ($false | (~ spl18_49 | ~ spl18_51)), inference(subsumption_resolution, [], [f2974, f188])).
fof(f2974, plain, ((e0 = e2) | (~ spl18_49 | ~ spl18_51)), inference(forward_demodulation, [], [f576, f568])).
fof(f2973, plain, (~ spl18_26 | ~ spl18_58), inference(avatar_split_clause, [], [f2969, f604, f468])).
fof(f2969, plain, (~ (e1 = op(e2, e1)) | ~ spl18_58), inference(backward_demodulation, [], [f146, f606])).
fof(f2968, plain, (~ spl18_32 | ~ spl18_41 | ~ spl18_63 | spl18_101), inference(avatar_contradiction_clause, [], [f2967])).
fof(f2967, plain, ($false | (~ spl18_32 | ~ spl18_41 | ~ spl18_63 | spl18_101)), inference(subsumption_resolution, [], [f2960, f495])).
fof(f495, plain, ((e3 = op(e2, e0)) | ~ spl18_32), inference(avatar_component_clause, [], [f493])).
fof(f2960, plain, (~ (e3 = op(e2, e0)) | (~ spl18_41 | ~ spl18_63 | spl18_101)), inference(backward_demodulation, [], [f2948, f627])).
fof(f2948, plain, (~ (e3 = op(op(e0, e0), e0)) | (~ spl18_41 | spl18_101)), inference(forward_demodulation, [], [f937, f534])).
fof(f937, plain, (~ (e3 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | spl18_101), inference(avatar_component_clause, [], [f935])).
fof(f935, plain, (spl18_101 <=> (e3 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl18_101])])).
fof(f2938, plain, (~ spl18_59 | ~ spl18_2 | ~ spl18_41 | spl18_109), inference(avatar_split_clause, [], [f2872, f971, f532, f366, f608])).
fof(f2872, plain, (~ (e2 = op(e0, e1)) | (~ spl18_2 | ~ spl18_41 | spl18_109)), inference(backward_demodulation, [], [f2829, f534])).
fof(f2829, plain, (~ (e2 = op(op(e1, e1), e1)) | (~ spl18_2 | spl18_109)), inference(backward_demodulation, [], [f973, f368])).
fof(f2933, plain, (~ spl18_48 | ~ spl18_21 | spl18_105 | ~ spl18_106), inference(avatar_split_clause, [], [f2900, f957, f953, f447, f561])).
fof(f2900, plain, (~ (e3 = op(e1, e0)) | (~ spl18_21 | spl18_105 | ~ spl18_106)), inference(backward_demodulation, [], [f2853, f449])).
fof(f2853, plain, (~ (e3 = op(e1, op(e2, e2))) | (spl18_105 | ~ spl18_106)), inference(forward_demodulation, [], [f955, f958])).
fof(f2927, plain, (~ spl18_28 | ~ spl18_32), inference(avatar_split_clause, [], [f2926, f493, f476])).
fof(f2926, plain, (~ (e3 = op(e2, e1)) | ~ spl18_32), inference(forward_demodulation, [], [f175, f495])).
fof(f2922, plain, (~ spl18_20 | ~ spl18_32), inference(avatar_split_clause, [], [f2921, f493, f442])).
fof(f2921, plain, (~ (e3 = op(e2, e3)) | ~ spl18_32), inference(forward_demodulation, [], [f177, f495])).
fof(f2920, plain, (~ spl18_20 | ~ spl18_36), inference(avatar_split_clause, [], [f2919, f510, f442])).
fof(f2919, plain, (~ (e3 = op(e2, e3)) | ~ spl18_36), inference(forward_demodulation, [], [f160, f512])).
fof(f512, plain, ((e3 = op(e1, e3)) | ~ spl18_36), inference(avatar_component_clause, [], [f510])).
fof(f2914, plain, (~ spl18_8 | ~ spl18_12), inference(avatar_split_clause, [], [f2913, f408, f391])).
fof(f2913, plain, (~ (e3 = op(e3, e2)) | ~ spl18_12), inference(forward_demodulation, [], [f184, f410])).
fof(f2910, plain, (~ spl18_8 | ~ spl18_56), inference(avatar_split_clause, [], [f2909, f595, f391])).
fof(f2909, plain, (~ (e3 = op(e3, e2)) | ~ spl18_56), inference(forward_demodulation, [], [f153, f597])).
fof(f597, plain, ((e3 = op(e0, e2)) | ~ spl18_56), inference(avatar_component_clause, [], [f595])).
fof(f2907, plain, (~ spl18_9 | ~ spl18_13), inference(avatar_split_clause, [], [f2902, f413, f396])).
fof(f2902, plain, (~ (e0 = op(e3, e1)) | ~ spl18_13), inference(backward_demodulation, [], [f181, f415])).
fof(f2891, plain, (~ spl18_24 | ~ spl18_32), inference(avatar_split_clause, [], [f2889, f493, f459])).
fof(f2889, plain, (~ (e3 = op(e2, e2)) | ~ spl18_32), inference(backward_demodulation, [], [f176, f495])).
fof(f2890, plain, (~ spl18_16 | ~ spl18_32), inference(avatar_split_clause, [], [f2888, f493, f425])).
fof(f2888, plain, (~ (e3 = op(e3, e0)) | ~ spl18_32), inference(backward_demodulation, [], [f144, f495])).
fof(f2874, plain, (~ spl18_9 | ~ spl18_41), inference(avatar_split_clause, [], [f2862, f532, f396])).
fof(f2862, plain, (~ (e0 = op(e3, e1)) | ~ spl18_41), inference(backward_demodulation, [], [f149, f534])).
fof(f2855, plain, (spl18_12 | ~ spl18_2 | ~ spl18_95), inference(avatar_split_clause, [], [f2854, f906, f366, f408])).
fof(f2854, plain, ((e3 = op(e3, e1)) | (~ spl18_2 | ~ spl18_95)), inference(forward_demodulation, [], [f908, f368])).
fof(f2852, plain, (~ spl18_44 | ~ spl18_62 | spl18_112), inference(avatar_split_clause, [], [f2851, f984, f621, f544])).
fof(f2851, plain, (~ (e3 = op(e1, e1)) | (~ spl18_62 | spl18_112)), inference(forward_demodulation, [], [f986, f623])).
fof(f2849, plain, (~ spl18_43 | ~ spl18_2 | spl18_108), inference(avatar_split_clause, [], [f2828, f966, f366, f540])).
fof(f2828, plain, (~ (e2 = op(e1, e1)) | (~ spl18_2 | spl18_108)), inference(backward_demodulation, [], [f968, f368])).
fof(f2847, plain, (~ spl18_41 | ~ spl18_2 | spl18_118), inference(avatar_split_clause, [], [f2832, f1013, f366, f532])).
fof(f2832, plain, (~ (e0 = op(e1, e1)) | (~ spl18_2 | spl18_118)), inference(backward_demodulation, [], [f1015, f368])).
fof(f2845, plain, (~ spl18_30 | ~ spl18_62), inference(avatar_split_clause, [], [f2734, f621, f485])).
fof(f2734, plain, (~ (e1 = op(e2, e0)) | ~ spl18_62), inference(backward_demodulation, [], [f140, f623])).
fof(f140, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f2823, plain, (~ spl18_1 | ~ spl18_9), inference(avatar_split_clause, [], [f2818, f396, f362])).
fof(f2818, plain, (~ (e0 = op(e3, e3)) | ~ spl18_9), inference(backward_demodulation, [], [f185, f398])).
fof(f398, plain, ((e0 = op(e3, e1)) | ~ spl18_9), inference(avatar_component_clause, [], [f396])).
fof(f2782, plain, (~ spl18_12 | ~ spl18_44), inference(avatar_split_clause, [], [f2775, f544, f408])).
fof(f2775, plain, (~ (e3 = op(e3, e1)) | ~ spl18_44), inference(backward_demodulation, [], [f149, f546])).
fof(f2774, plain, (~ spl18_43 | ~ spl18_47), inference(avatar_split_clause, [], [f2771, f557, f540])).
fof(f2771, plain, (~ (e2 = op(e1, e1)) | ~ spl18_47), inference(backward_demodulation, [], [f169, f559])).
fof(f2761, plain, (~ spl18_54 | ~ spl18_56), inference(avatar_contradiction_clause, [], [f2760])).
fof(f2760, plain, ($false | (~ spl18_54 | ~ spl18_56)), inference(subsumption_resolution, [], [f2759, f191])).
fof(f2759, plain, ((e1 = e3) | (~ spl18_54 | ~ spl18_56)), inference(backward_demodulation, [], [f597, f589])).
fof(f2757, plain, (spl18_54 | ~ spl18_59 | ~ spl18_93), inference(avatar_split_clause, [], [f2755, f887, f608, f587])).
fof(f2755, plain, ((e1 = op(e0, e2)) | (~ spl18_59 | ~ spl18_93)), inference(backward_demodulation, [], [f889, f610])).
fof(f2756, plain, (~ spl18_43 | ~ spl18_59), inference(avatar_split_clause, [], [f2751, f608, f540])).
fof(f2751, plain, (~ (e2 = op(e1, e1)) | ~ spl18_59), inference(backward_demodulation, [], [f145, f610])).
fof(f2748, plain, (~ spl18_43 | ~ spl18_62 | spl18_114), inference(avatar_split_clause, [], [f2742, f993, f621, f540])).
fof(f2742, plain, (~ (e2 = op(e1, e1)) | (~ spl18_62 | spl18_114)), inference(backward_demodulation, [], [f995, f623])).
fof(f2746, plain, (~ spl18_57 | ~ spl18_62 | spl18_94), inference(avatar_split_clause, [], [f2740, f892, f621, f600])).
fof(f2740, plain, (~ (e0 = op(e0, e1)) | (~ spl18_62 | spl18_94)), inference(backward_demodulation, [], [f893, f623])).
fof(f893, plain, (~ (e0 = op(e0, op(e0, e0))) | spl18_94), inference(avatar_component_clause, [], [f892])).
fof(f2732, plain, (spl18_24 | ~ spl18_19 | ~ spl18_81), inference(avatar_split_clause, [], [f2731, f823, f438, f459])).
fof(f2731, plain, ((e3 = op(e2, e2)) | (~ spl18_19 | ~ spl18_81)), inference(forward_demodulation, [], [f825, f440])).
fof(f440, plain, ((e2 = op(e2, e3)) | ~ spl18_19), inference(avatar_component_clause, [], [f438])).
fof(f2716, plain, (~ spl18_42 | ~ spl18_26), inference(avatar_split_clause, [], [f2715, f468, f536])).
fof(f2715, plain, (~ (e1 = op(e1, e1)) | ~ spl18_26), inference(forward_demodulation, [], [f148, f470])).
fof(f2712, plain, (~ spl18_44 | ~ spl18_36), inference(avatar_split_clause, [], [f2711, f510, f544])).
fof(f2711, plain, (~ (e3 = op(e1, e1)) | ~ spl18_36), inference(forward_demodulation, [], [f173, f512])).
fof(f2706, plain, (~ spl18_24 | ~ spl18_56), inference(avatar_split_clause, [], [f2705, f595, f459])).
fof(f2705, plain, (~ (e3 = op(e2, e2)) | ~ spl18_56), inference(forward_demodulation, [], [f152, f597])).
fof(f2689, plain, (~ spl18_2 | ~ spl18_14), inference(avatar_split_clause, [], [f2685, f417, f366])).
fof(f2685, plain, (~ (e1 = op(e3, e3)) | ~ spl18_14), inference(backward_demodulation, [], [f183, f419])).
fof(f2658, plain, (~ spl18_50 | ~ spl18_57 | ~ spl18_91), inference(avatar_contradiction_clause, [], [f2657])).
fof(f2657, plain, ($false | (~ spl18_50 | ~ spl18_57 | ~ spl18_91)), inference(subsumption_resolution, [], [f2656, f189])).
fof(f2656, plain, ((e0 = e3) | (~ spl18_50 | ~ spl18_57 | ~ spl18_91)), inference(forward_demodulation, [], [f2654, f602])).
fof(f2654, plain, ((e3 = op(e0, e1)) | (~ spl18_50 | ~ spl18_91)), inference(backward_demodulation, [], [f879, f572])).
fof(f2655, plain, (~ spl18_2 | ~ spl18_50), inference(avatar_split_clause, [], [f2651, f570, f366])).
fof(f2651, plain, (~ (e1 = op(e3, e3)) | ~ spl18_50), inference(backward_demodulation, [], [f159, f572])).
fof(f2649, plain, (~ spl18_52 | ~ spl18_56), inference(avatar_split_clause, [], [f2645, f595, f578])).
fof(f2645, plain, (~ (e3 = op(e0, e3)) | ~ spl18_56), inference(backward_demodulation, [], [f168, f597])).
fof(f2648, plain, (~ spl18_40 | ~ spl18_56), inference(avatar_split_clause, [], [f2644, f595, f527])).
fof(f2644, plain, (~ (e3 = op(e1, e2)) | ~ spl18_56), inference(backward_demodulation, [], [f151, f597])).
fof(f2647, plain, (spl18_51 | ~ spl18_56 | ~ spl18_92), inference(avatar_contradiction_clause, [], [f2646])).
fof(f2646, plain, ($false | (spl18_51 | ~ spl18_56 | ~ spl18_92)), inference(subsumption_resolution, [], [f2643, f575])).
fof(f575, plain, (~ (e2 = op(e0, e3)) | spl18_51), inference(avatar_component_clause, [], [f574])).
fof(f2643, plain, ((e2 = op(e0, e3)) | (~ spl18_56 | ~ spl18_92)), inference(backward_demodulation, [], [f884, f597])).
fof(f884, plain, ((e2 = op(e0, op(e0, e2))) | ~ spl18_92), inference(avatar_component_clause, [], [f882])).
fof(f2623, plain, (spl18_53 | ~ spl18_63 | ~ spl18_94), inference(avatar_contradiction_clause, [], [f2622])).
fof(f2622, plain, ($false | (spl18_53 | ~ spl18_63 | ~ spl18_94)), inference(subsumption_resolution, [], [f2612, f584])).
fof(f584, plain, (~ (e0 = op(e0, e2)) | spl18_53), inference(avatar_component_clause, [], [f583])).
fof(f2612, plain, ((e0 = op(e0, e2)) | (~ spl18_63 | ~ spl18_94)), inference(backward_demodulation, [], [f894, f627])).
fof(f2621, plain, (~ spl18_55 | ~ spl18_63), inference(avatar_split_clause, [], [f2609, f625, f591])).
fof(f2609, plain, (~ (e2 = op(e0, e2)) | ~ spl18_63), inference(backward_demodulation, [], [f164, f627])).
fof(f2620, plain, (~ spl18_31 | ~ spl18_63), inference(avatar_split_clause, [], [f2606, f625, f489])).
fof(f2606, plain, (~ (e2 = op(e2, e0)) | ~ spl18_63), inference(backward_demodulation, [], [f140, f627])).
fof(f2604, plain, (~ spl18_31 | ~ spl18_21 | spl18_82), inference(avatar_split_clause, [], [f2603, f828, f447, f489])).
fof(f2603, plain, (~ (e2 = op(e2, e0)) | (~ spl18_21 | spl18_82)), inference(forward_demodulation, [], [f829, f449])).
fof(f2602, plain, (spl18_38 | ~ spl18_43 | ~ spl18_88), inference(avatar_split_clause, [], [f2601, f860, f540, f519])).
fof(f2601, plain, ((e1 = op(e1, e2)) | (~ spl18_43 | ~ spl18_88)), inference(forward_demodulation, [], [f862, f542])).
fof(f2600, plain, (~ spl18_56 | ~ spl18_21 | ~ spl18_43 | spl18_101), inference(avatar_split_clause, [], [f2599, f935, f540, f447, f595])).
fof(f2599, plain, (~ (e3 = op(e0, e2)) | (~ spl18_21 | ~ spl18_43 | spl18_101)), inference(forward_demodulation, [], [f2598, f449])).
fof(f2598, plain, (~ (e3 = op(op(e2, e2), e2)) | (~ spl18_43 | spl18_101)), inference(forward_demodulation, [], [f937, f542])).
fof(f2592, plain, (~ spl18_35 | ~ spl18_43), inference(avatar_split_clause, [], [f2591, f540, f506])).
fof(f2591, plain, (~ (e2 = op(e1, e3)) | ~ spl18_43), inference(forward_demodulation, [], [f173, f542])).
fof(f2585, plain, (~ spl18_64 | ~ spl18_21 | spl18_104), inference(avatar_split_clause, [], [f2584, f948, f447, f629])).
fof(f2584, plain, (~ (op(e0, e0) = e3) | (~ spl18_21 | spl18_104)), inference(forward_demodulation, [], [f950, f449])).
fof(f950, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | spl18_104), inference(avatar_component_clause, [], [f948])).
fof(f2582, plain, (~ spl18_62 | ~ spl18_21 | spl18_106), inference(avatar_split_clause, [], [f2581, f957, f447, f621])).
fof(f2581, plain, (~ (op(e0, e0) = e1) | (~ spl18_21 | spl18_106)), inference(forward_demodulation, [], [f959, f449])).
fof(f959, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | spl18_106), inference(avatar_component_clause, [], [f957])).
fof(f2575, plain, (~ spl18_11 | ~ spl18_43), inference(avatar_split_clause, [], [f2574, f540, f404])).
fof(f2574, plain, (~ (e2 = op(e3, e1)) | ~ spl18_43), inference(forward_demodulation, [], [f149, f542])).
fof(f2570, plain, (~ spl18_21 | ~ spl18_22), inference(avatar_contradiction_clause, [], [f2569])).
fof(f2569, plain, ($false | (~ spl18_21 | ~ spl18_22)), inference(subsumption_resolution, [], [f2568, f187])).
fof(f2568, plain, ((e0 = e1) | (~ spl18_21 | ~ spl18_22)), inference(backward_demodulation, [], [f453, f449])).
fof(f2550, plain, (spl18_18 | ~ spl18_28 | ~ spl18_83), inference(avatar_contradiction_clause, [], [f2549])).
fof(f2549, plain, ($false | (spl18_18 | ~ spl18_28 | ~ spl18_83)), inference(subsumption_resolution, [], [f2547, f435])).
fof(f435, plain, (~ (e1 = op(e2, e3)) | spl18_18), inference(avatar_component_clause, [], [f434])).
fof(f2547, plain, ((e1 = op(e2, e3)) | (~ spl18_28 | ~ spl18_83)), inference(backward_demodulation, [], [f835, f478])).
fof(f2545, plain, (spl18_21 | ~ spl18_31 | ~ spl18_84), inference(avatar_split_clause, [], [f2544, f838, f489, f447])).
fof(f2544, plain, ((e0 = op(e2, e2)) | (~ spl18_31 | ~ spl18_84)), inference(backward_demodulation, [], [f840, f491])).
fof(f840, plain, ((e0 = op(e2, op(e2, e0))) | ~ spl18_84), inference(avatar_component_clause, [], [f838])).
fof(f2533, plain, (~ spl18_38 | ~ spl18_43 | spl18_88), inference(avatar_split_clause, [], [f2525, f860, f540, f519])).
fof(f2525, plain, (~ (e1 = op(e1, e2)) | (~ spl18_43 | spl18_88)), inference(backward_demodulation, [], [f861, f542])).
fof(f2495, plain, (~ spl18_53 | ~ spl18_61), inference(avatar_split_clause, [], [f2484, f617, f583])).
fof(f2484, plain, (~ (e0 = op(e0, e2)) | ~ spl18_61), inference(backward_demodulation, [], [f164, f619])).
fof(f619, plain, ((e0 = op(e0, e0)) | ~ spl18_61), inference(avatar_component_clause, [], [f617])).
fof(f2494, plain, (~ spl18_29 | ~ spl18_61), inference(avatar_split_clause, [], [f2482, f617, f481])).
fof(f2482, plain, (~ (e0 = op(e2, e0)) | ~ spl18_61), inference(backward_demodulation, [], [f140, f619])).
fof(f2468, plain, (~ spl18_55 | ~ spl18_7), inference(avatar_split_clause, [], [f2467, f387, f591])).
fof(f2467, plain, (~ (e2 = op(e0, e2)) | ~ spl18_7), inference(forward_demodulation, [], [f153, f389])).
fof(f389, plain, ((e2 = op(e3, e2)) | ~ spl18_7), inference(avatar_component_clause, [], [f387])).
fof(f2464, plain, (~ spl18_48 | ~ spl18_33 | spl18_86), inference(avatar_split_clause, [], [f2437, f850, f498, f561])).
fof(f2437, plain, (~ (e3 = op(e1, e0)) | (~ spl18_33 | spl18_86)), inference(forward_demodulation, [], [f851, f500])).
fof(f500, plain, ((e0 = op(e1, e3)) | ~ spl18_33), inference(avatar_component_clause, [], [f498])).
fof(f2462, plain, (~ spl18_31 | ~ spl18_19), inference(avatar_split_clause, [], [f2018, f438, f489])).
fof(f2018, plain, (~ (e2 = op(e2, e0)) | ~ spl18_19), inference(forward_demodulation, [], [f177, f440])).
fof(f2436, plain, (~ spl18_43 | ~ spl18_38 | spl18_87), inference(avatar_split_clause, [], [f2435, f855, f519, f540])).
fof(f855, plain, (spl18_87 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl18_87])])).
fof(f2435, plain, (~ (e2 = op(e1, e1)) | (~ spl18_38 | spl18_87)), inference(forward_demodulation, [], [f856, f521])).
fof(f856, plain, (~ (e2 = op(e1, op(e1, e2))) | spl18_87), inference(avatar_component_clause, [], [f855])).
fof(f2434, plain, (~ spl18_63 | ~ spl18_53 | spl18_92), inference(avatar_split_clause, [], [f2433, f882, f583, f625])).
fof(f2433, plain, (~ (op(e0, e0) = e2) | (~ spl18_53 | spl18_92)), inference(forward_demodulation, [], [f883, f585])).
fof(f2431, plain, (~ spl18_2 | ~ spl18_24 | ~ spl18_33 | spl18_115), inference(avatar_contradiction_clause, [], [f2430])).
fof(f2430, plain, ($false | (~ spl18_2 | ~ spl18_24 | ~ spl18_33 | spl18_115)), inference(subsumption_resolution, [], [f2429, f500])).
fof(f2429, plain, (~ (e0 = op(e1, e3)) | (~ spl18_2 | ~ spl18_24 | spl18_115)), inference(forward_demodulation, [], [f2428, f368])).
fof(f2428, plain, (~ (e0 = op(op(e3, e3), e3)) | (~ spl18_24 | spl18_115)), inference(forward_demodulation, [], [f1000, f461])).
fof(f2427, plain, (~ spl18_2 | ~ spl18_24 | spl18_106), inference(avatar_contradiction_clause, [], [f2426])).
fof(f2426, plain, ($false | (~ spl18_2 | ~ spl18_24 | spl18_106)), inference(subsumption_resolution, [], [f2425, f368])).
fof(f2425, plain, (~ (e1 = op(e3, e3)) | (~ spl18_24 | spl18_106)), inference(forward_demodulation, [], [f959, f461])).
fof(f2402, plain, (~ spl18_60 | ~ spl18_44), inference(avatar_split_clause, [], [f2401, f544, f612])).
fof(f2401, plain, (~ (e3 = op(e0, e1)) | ~ spl18_44), inference(forward_demodulation, [], [f145, f546])).
fof(f2366, plain, (~ spl18_29 | ~ spl18_30), inference(avatar_contradiction_clause, [], [f2365])).
fof(f2365, plain, ($false | (~ spl18_29 | ~ spl18_30)), inference(subsumption_resolution, [], [f2364, f187])).
fof(f2364, plain, ((e0 = e1) | (~ spl18_29 | ~ spl18_30)), inference(forward_demodulation, [], [f487, f483])).
fof(f2363, plain, (~ spl18_43 | ~ spl18_44), inference(avatar_contradiction_clause, [], [f2362])).
fof(f2362, plain, ($false | (~ spl18_43 | ~ spl18_44)), inference(subsumption_resolution, [], [f2361, f192])).
fof(f2361, plain, ((e2 = e3) | (~ spl18_43 | ~ spl18_44)), inference(forward_demodulation, [], [f546, f542])).
fof(f2338, plain, (~ spl18_50 | ~ spl18_24 | spl18_103 | ~ spl18_116), inference(avatar_split_clause, [], [f2307, f1003, f944, f459, f570])).
fof(f2307, plain, (~ (e1 = op(e0, e3)) | (~ spl18_24 | spl18_103 | ~ spl18_116)), inference(backward_demodulation, [], [f2253, f461])).
fof(f2253, plain, (~ (e1 = op(e0, op(e2, e2))) | (spl18_103 | ~ spl18_116)), inference(forward_demodulation, [], [f946, f1004])).
fof(f2334, plain, (spl18_48 | ~ spl18_33 | ~ spl18_86), inference(avatar_split_clause, [], [f2284, f850, f498, f561])).
fof(f2284, plain, ((e3 = op(e1, e0)) | (~ spl18_33 | ~ spl18_86)), inference(backward_demodulation, [], [f852, f500])).
fof(f2322, plain, (~ spl18_1 | ~ spl18_33), inference(avatar_split_clause, [], [f2321, f498, f362])).
fof(f2321, plain, (~ (e0 = op(e3, e3)) | ~ spl18_33), inference(forward_demodulation, [], [f161, f500])).
fof(f2318, plain, (~ spl18_9 | ~ spl18_12), inference(avatar_contradiction_clause, [], [f2317])).
fof(f2317, plain, ($false | (~ spl18_9 | ~ spl18_12)), inference(subsumption_resolution, [], [f2314, f189])).
fof(f2314, plain, ((e0 = e3) | (~ spl18_9 | ~ spl18_12)), inference(backward_demodulation, [], [f410, f398])).
fof(f2298, plain, (~ spl18_26 | spl18_83), inference(avatar_contradiction_clause, [], [f2297])).
fof(f2297, plain, ($false | (~ spl18_26 | spl18_83)), inference(subsumption_resolution, [], [f2295, f470])).
fof(f2295, plain, (~ (e1 = op(e2, e1)) | (~ spl18_26 | spl18_83)), inference(backward_demodulation, [], [f834, f470])).
fof(f834, plain, (~ (e1 = op(e2, op(e2, e1))) | spl18_83), inference(avatar_component_clause, [], [f833])).
fof(f2293, plain, (~ spl18_25 | ~ spl18_29), inference(avatar_split_clause, [], [f2289, f481, f464])).
fof(f2289, plain, (~ (e0 = op(e2, e1)) | ~ spl18_29), inference(backward_demodulation, [], [f175, f483])).
fof(f2274, plain, (spl18_33 | ~ spl18_48 | ~ spl18_89), inference(avatar_split_clause, [], [f2271, f865, f561, f498])).
fof(f2271, plain, ((e0 = op(e1, e3)) | (~ spl18_48 | ~ spl18_89)), inference(backward_demodulation, [], [f867, f563])).
fof(f2268, plain, (~ spl18_50 | ~ spl18_52), inference(avatar_contradiction_clause, [], [f2267])).
fof(f2267, plain, ($false | (~ spl18_50 | ~ spl18_52)), inference(subsumption_resolution, [], [f2266, f191])).
fof(f2266, plain, ((e1 = e3) | (~ spl18_50 | ~ spl18_52)), inference(backward_demodulation, [], [f580, f572])).
fof(f2264, plain, (spl18_50 | ~ spl18_60 | ~ spl18_93), inference(avatar_split_clause, [], [f2263, f887, f612, f570])).
fof(f2263, plain, ((e1 = op(e0, e3)) | (~ spl18_60 | ~ spl18_93)), inference(backward_demodulation, [], [f889, f614])).
fof(f2230, plain, (~ spl18_23 | ~ spl18_7), inference(avatar_split_clause, [], [f2229, f387, f455])).
fof(f2229, plain, (~ (e2 = op(e2, e2)) | ~ spl18_7), inference(forward_demodulation, [], [f156, f389])).
fof(f2228, plain, (~ spl18_23 | ~ spl18_19), inference(avatar_split_clause, [], [f2227, f438, f455])).
fof(f2227, plain, (~ (e2 = op(e2, e2)) | ~ spl18_19), inference(forward_demodulation, [], [f180, f440])).
fof(f2203, plain, (spl18_2 | ~ spl18_12 | ~ spl18_97), inference(avatar_split_clause, [], [f2202, f916, f408, f366])).
fof(f916, plain, (spl18_97 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl18_97])])).
fof(f2202, plain, ((e1 = op(e3, e3)) | (~ spl18_12 | ~ spl18_97)), inference(backward_demodulation, [], [f918, f410])).
fof(f918, plain, ((e1 = op(e3, op(e3, e1))) | ~ spl18_97), inference(avatar_component_clause, [], [f916])).
fof(f2193, plain, (~ spl18_9 | ~ spl18_25), inference(avatar_split_clause, [], [f2190, f464, f396])).
fof(f2190, plain, (~ (e0 = op(e3, e1)) | ~ spl18_25), inference(backward_demodulation, [], [f150, f466])).
fof(f2158, plain, (~ spl18_22 | ~ spl18_40 | ~ spl18_63 | spl18_113), inference(avatar_contradiction_clause, [], [f2157])).
fof(f2157, plain, ($false | (~ spl18_22 | ~ spl18_40 | ~ spl18_63 | spl18_113)), inference(subsumption_resolution, [], [f2156, f529])).
fof(f2156, plain, (~ (e3 = op(e1, e2)) | (~ spl18_22 | ~ spl18_63 | spl18_113)), inference(forward_demodulation, [], [f2155, f453])).
fof(f2155, plain, (~ (e3 = op(op(e2, e2), e2)) | (~ spl18_63 | spl18_113)), inference(forward_demodulation, [], [f991, f627])).
fof(f2139, plain, (~ spl18_2 | ~ spl18_9 | ~ spl18_95), inference(avatar_contradiction_clause, [], [f2138])).
fof(f2138, plain, ($false | (~ spl18_2 | ~ spl18_9 | ~ spl18_95)), inference(subsumption_resolution, [], [f2137, f189])).
fof(f2137, plain, ((e0 = e3) | (~ spl18_2 | ~ spl18_9 | ~ spl18_95)), inference(forward_demodulation, [], [f2131, f398])).
fof(f2131, plain, ((e3 = op(e3, e1)) | (~ spl18_2 | ~ spl18_95)), inference(backward_demodulation, [], [f908, f368])).
fof(f2128, plain, (~ spl18_9 | spl18_14 | ~ spl18_97), inference(avatar_contradiction_clause, [], [f2127])).
fof(f2127, plain, ($false | (~ spl18_9 | spl18_14 | ~ spl18_97)), inference(subsumption_resolution, [], [f2126, f418])).
fof(f418, plain, (~ (e1 = op(e3, e0)) | spl18_14), inference(avatar_component_clause, [], [f417])).
fof(f2126, plain, ((e1 = op(e3, e0)) | (~ spl18_9 | ~ spl18_97)), inference(backward_demodulation, [], [f918, f398])).
fof(f2113, plain, (~ spl18_29 | spl18_84), inference(avatar_contradiction_clause, [], [f2112])).
fof(f2112, plain, ($false | (~ spl18_29 | spl18_84)), inference(subsumption_resolution, [], [f2110, f483])).
fof(f2110, plain, (~ (e0 = op(e2, e0)) | (~ spl18_29 | spl18_84)), inference(backward_demodulation, [], [f839, f483])).
fof(f2076, plain, (~ spl18_52 | spl18_91), inference(avatar_contradiction_clause, [], [f2075])).
fof(f2075, plain, ($false | (~ spl18_52 | spl18_91)), inference(subsumption_resolution, [], [f2073, f580])).
fof(f2073, plain, (~ (e3 = op(e0, e3)) | (~ spl18_52 | spl18_91)), inference(backward_demodulation, [], [f878, f580])).
fof(f878, plain, (~ (e3 = op(e0, op(e0, e3))) | spl18_91), inference(avatar_component_clause, [], [f877])).
fof(f2070, plain, (~ spl18_50 | ~ spl18_58), inference(avatar_split_clause, [], [f2065, f604, f570])).
fof(f2065, plain, (~ (e1 = op(e0, e3)) | ~ spl18_58), inference(backward_demodulation, [], [f167, f606])).
fof(f2039, plain, (spl18_63 | ~ spl18_53 | ~ spl18_92), inference(avatar_split_clause, [], [f2038, f882, f583, f625])).
fof(f2038, plain, ((op(e0, e0) = e2) | (~ spl18_53 | ~ spl18_92)), inference(forward_demodulation, [], [f884, f585])).
fof(f2033, plain, (~ spl18_51 | ~ spl18_19), inference(avatar_split_clause, [], [f2032, f438, f574])).
fof(f2032, plain, (~ (e2 = op(e0, e3)) | ~ spl18_19), inference(forward_demodulation, [], [f158, f440])).
fof(f2011, plain, (~ spl18_15 | ~ spl18_7), inference(avatar_split_clause, [], [f2010, f387, f421])).
fof(f2010, plain, (~ (e2 = op(e3, e0)) | ~ spl18_7), inference(forward_demodulation, [], [f182, f389])).
fof(f2009, plain, (~ spl18_11 | ~ spl18_7), inference(avatar_split_clause, [], [f2008, f387, f404])).
fof(f2008, plain, (~ (e2 = op(e3, e1)) | ~ spl18_7), inference(forward_demodulation, [], [f184, f389])).
fof(f2003, plain, (~ spl18_10 | ~ spl18_12), inference(avatar_contradiction_clause, [], [f2002])).
fof(f2002, plain, ($false | (~ spl18_10 | ~ spl18_12)), inference(subsumption_resolution, [], [f2001, f191])).
fof(f2001, plain, ((e1 = e3) | (~ spl18_10 | ~ spl18_12)), inference(forward_demodulation, [], [f410, f402])).
fof(f2000, plain, (~ spl18_17 | ~ spl18_19), inference(avatar_contradiction_clause, [], [f1999])).
fof(f1999, plain, ($false | (~ spl18_17 | ~ spl18_19)), inference(subsumption_resolution, [], [f1997, f188])).
fof(f1997, plain, ((e0 = e2) | (~ spl18_17 | ~ spl18_19)), inference(backward_demodulation, [], [f440, f432])).
fof(f1996, plain, (~ spl18_19 | ~ spl18_22 | ~ spl18_81), inference(avatar_contradiction_clause, [], [f1995])).
fof(f1995, plain, ($false | (~ spl18_19 | ~ spl18_22 | ~ spl18_81)), inference(subsumption_resolution, [], [f1994, f191])).
fof(f1994, plain, ((e1 = e3) | (~ spl18_19 | ~ spl18_22 | ~ spl18_81)), inference(forward_demodulation, [], [f1993, f453])).
fof(f1993, plain, ((e3 = op(e2, e2)) | (~ spl18_19 | ~ spl18_81)), inference(backward_demodulation, [], [f825, f440])).
fof(f1991, plain, (~ spl18_22 | ~ spl18_24), inference(avatar_contradiction_clause, [], [f1990])).
fof(f1990, plain, ($false | (~ spl18_22 | ~ spl18_24)), inference(subsumption_resolution, [], [f1989, f191])).
fof(f1989, plain, ((e1 = e3) | (~ spl18_22 | ~ spl18_24)), inference(forward_demodulation, [], [f461, f453])).
fof(f1988, plain, (~ spl18_27 | ~ spl18_28), inference(avatar_contradiction_clause, [], [f1987])).
fof(f1987, plain, ($false | (~ spl18_27 | ~ spl18_28)), inference(subsumption_resolution, [], [f1986, f192])).
fof(f1986, plain, ((e2 = e3) | (~ spl18_27 | ~ spl18_28)), inference(backward_demodulation, [], [f478, f474])).
fof(f1964, plain, (~ spl18_33 | ~ spl18_41), inference(avatar_split_clause, [], [f1953, f532, f498])).
fof(f1953, plain, (~ (e0 = op(e1, e3)) | ~ spl18_41), inference(backward_demodulation, [], [f173, f534])).
fof(f1951, plain, (~ spl18_46 | ~ spl18_47), inference(avatar_contradiction_clause, [], [f1950])).
fof(f1950, plain, ($false | (~ spl18_46 | ~ spl18_47)), inference(subsumption_resolution, [], [f1949, f190])).
fof(f190, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f1949, plain, ((e1 = e2) | (~ spl18_46 | ~ spl18_47)), inference(backward_demodulation, [], [f559, f555])).
fof(f1948, plain, (~ spl18_31 | ~ spl18_47), inference(avatar_split_clause, [], [f1945, f557, f489])).
fof(f1945, plain, (~ (e2 = op(e2, e0)) | ~ spl18_47), inference(backward_demodulation, [], [f142, f559])).
fof(f1947, plain, (spl18_37 | ~ spl18_47 | ~ spl18_89), inference(avatar_contradiction_clause, [], [f1946])).
fof(f1946, plain, ($false | (spl18_37 | ~ spl18_47 | ~ spl18_89)), inference(subsumption_resolution, [], [f1942, f516])).
fof(f516, plain, (~ (e0 = op(e1, e2)) | spl18_37), inference(avatar_component_clause, [], [f515])).
fof(f1942, plain, ((e0 = op(e1, e2)) | (~ spl18_47 | ~ spl18_89)), inference(backward_demodulation, [], [f867, f559])).
fof(f1938, plain, (~ spl18_48 | ~ spl18_64), inference(avatar_split_clause, [], [f1931, f629, f561])).
fof(f1931, plain, (~ (e3 = op(e1, e0)) | ~ spl18_64), inference(backward_demodulation, [], [f139, f631])).
fof(f1904, plain, (~ spl18_41 | ~ spl18_22 | spl18_116), inference(avatar_split_clause, [], [f1524, f1003, f451, f532])).
fof(f1524, plain, (~ (e0 = op(e1, e1)) | (~ spl18_22 | spl18_116)), inference(backward_demodulation, [], [f1005, f453])).
fof(f1882, plain, (spl18_8 | ~ spl18_3 | ~ spl18_95), inference(avatar_split_clause, [], [f1875, f906, f370, f391])).
fof(f1875, plain, ((e3 = op(e3, e2)) | (~ spl18_3 | ~ spl18_95)), inference(backward_demodulation, [], [f908, f372])).
fof(f1861, plain, (~ spl18_39 | spl18_87), inference(avatar_contradiction_clause, [], [f1860])).
fof(f1860, plain, ($false | (~ spl18_39 | spl18_87)), inference(subsumption_resolution, [], [f1858, f525])).
fof(f1858, plain, (~ (e2 = op(e1, e2)) | (~ spl18_39 | spl18_87)), inference(backward_demodulation, [], [f856, f525])).
fof(f1859, plain, (~ spl18_35 | ~ spl18_39), inference(avatar_split_clause, [], [f1857, f523, f506])).
fof(f1857, plain, (~ (e2 = op(e1, e3)) | ~ spl18_39), inference(backward_demodulation, [], [f174, f525])).
fof(f1848, plain, (~ spl18_37 | ~ spl18_53), inference(avatar_split_clause, [], [f1844, f583, f515])).
fof(f1844, plain, (~ (e0 = op(e1, e2)) | ~ spl18_53), inference(backward_demodulation, [], [f151, f585])).
fof(f1847, plain, (~ spl18_49 | ~ spl18_53), inference(avatar_split_clause, [], [f1843, f583, f566])).
fof(f1843, plain, (~ (e0 = op(e0, e3)) | ~ spl18_53), inference(backward_demodulation, [], [f168, f585])).
fof(f1839, plain, (~ spl18_59 | ~ spl18_60), inference(avatar_contradiction_clause, [], [f1838])).
fof(f1838, plain, ($false | (~ spl18_59 | ~ spl18_60)), inference(subsumption_resolution, [], [f1837, f192])).
fof(f1837, plain, ((e2 = e3) | (~ spl18_59 | ~ spl18_60)), inference(backward_demodulation, [], [f614, f610])).
fof(f1815, plain, (~ spl18_64 | ~ spl18_60), inference(avatar_split_clause, [], [f1814, f612, f629])).
fof(f1814, plain, (~ (op(e0, e0) = e3) | ~ spl18_60), inference(forward_demodulation, [], [f163, f614])).
fof(f1803, plain, (~ spl18_38 | ~ spl18_42), inference(avatar_split_clause, [], [f1802, f536, f519])).
fof(f1802, plain, (~ (e1 = op(e1, e2)) | ~ spl18_42), inference(forward_demodulation, [], [f172, f538])).
fof(f538, plain, ((e1 = op(e1, e1)) | ~ spl18_42), inference(avatar_component_clause, [], [f536])).
fof(f1784, plain, (~ spl18_46 | ~ spl18_48), inference(avatar_contradiction_clause, [], [f1783])).
fof(f1783, plain, ($false | (~ spl18_46 | ~ spl18_48)), inference(subsumption_resolution, [], [f1782, f191])).
fof(f1782, plain, ((e1 = e3) | (~ spl18_46 | ~ spl18_48)), inference(forward_demodulation, [], [f563, f555])).
fof(f1781, plain, (~ spl18_58 | ~ spl18_60), inference(avatar_contradiction_clause, [], [f1780])).
fof(f1780, plain, ($false | (~ spl18_58 | ~ spl18_60)), inference(subsumption_resolution, [], [f1779, f191])).
fof(f1779, plain, ((e1 = e3) | (~ spl18_58 | ~ spl18_60)), inference(forward_demodulation, [], [f614, f606])).
fof(f1763, plain, (~ spl18_58 | spl18_93), inference(avatar_contradiction_clause, [], [f1762])).
fof(f1762, plain, ($false | (~ spl18_58 | spl18_93)), inference(subsumption_resolution, [], [f1761, f606])).
fof(f1761, plain, (~ (e1 = op(e0, e1)) | (~ spl18_58 | spl18_93)), inference(forward_demodulation, [], [f888, f606])).
fof(f888, plain, (~ (e1 = op(e0, op(e0, e1))) | spl18_93), inference(avatar_component_clause, [], [f887])).
fof(f1752, plain, (spl18_6 | ~ spl18_11 | ~ spl18_97), inference(avatar_contradiction_clause, [], [f1751])).
fof(f1751, plain, ($false | (spl18_6 | ~ spl18_11 | ~ spl18_97)), inference(subsumption_resolution, [], [f1750, f384])).
fof(f384, plain, (~ (e1 = op(e3, e2)) | spl18_6), inference(avatar_component_clause, [], [f383])).
fof(f1750, plain, ((e1 = op(e3, e2)) | (~ spl18_11 | ~ spl18_97)), inference(forward_demodulation, [], [f918, f406])).
fof(f1744, plain, (~ spl18_61 | ~ spl18_49), inference(avatar_split_clause, [], [f1743, f566, f617])).
fof(f1743, plain, (~ (e0 = op(e0, e0)) | ~ spl18_49), inference(forward_demodulation, [], [f165, f568])).
fof(f1742, plain, (~ spl18_61 | ~ spl18_13), inference(avatar_split_clause, [], [f1585, f413, f617])).
fof(f1585, plain, (~ (e0 = op(e0, e0)) | ~ spl18_13), inference(forward_demodulation, [], [f141, f415])).
fof(f1735, plain, (spl18_42 | ~ spl18_2 | ~ spl18_110), inference(avatar_split_clause, [], [f1721, f975, f366, f536])).
fof(f1721, plain, ((e1 = op(e1, e1)) | (~ spl18_2 | ~ spl18_110)), inference(backward_demodulation, [], [f976, f368])).
fof(f976, plain, ((e1 = op(op(e3, e3), op(e3, e3))) | ~ spl18_110), inference(avatar_component_clause, [], [f975])).
fof(f1732, plain, (~ spl18_2 | ~ spl18_35 | ~ spl18_64 | spl18_111), inference(avatar_contradiction_clause, [], [f1731])).
fof(f1731, plain, ($false | (~ spl18_2 | ~ spl18_35 | ~ spl18_64 | spl18_111)), inference(subsumption_resolution, [], [f1722, f508])).
fof(f1722, plain, (~ (e2 = op(e1, e3)) | (~ spl18_2 | ~ spl18_64 | spl18_111)), inference(backward_demodulation, [], [f1596, f368])).
fof(f1596, plain, (~ (e2 = op(op(e3, e3), e3)) | (~ spl18_64 | spl18_111)), inference(backward_demodulation, [], [f982, f631])).
fof(f1714, plain, (~ spl18_35 | ~ spl18_37 | ~ spl18_86), inference(avatar_contradiction_clause, [], [f1713])).
fof(f1713, plain, ($false | (~ spl18_35 | ~ spl18_37 | ~ spl18_86)), inference(subsumption_resolution, [], [f1712, f189])).
fof(f1712, plain, ((e0 = e3) | (~ spl18_35 | ~ spl18_37 | ~ spl18_86)), inference(forward_demodulation, [], [f1710, f517])).
fof(f1710, plain, ((e3 = op(e1, e2)) | (~ spl18_35 | ~ spl18_86)), inference(backward_demodulation, [], [f852, f508])).
fof(f1707, plain, (~ spl18_33 | ~ spl18_37), inference(avatar_split_clause, [], [f1702, f515, f498])).
fof(f1702, plain, (~ (e0 = op(e1, e3)) | ~ spl18_37), inference(backward_demodulation, [], [f174, f517])).
fof(f1706, plain, (~ spl18_37 | ~ spl18_46 | ~ spl18_87), inference(avatar_contradiction_clause, [], [f1705])).
fof(f1705, plain, ($false | (~ spl18_37 | ~ spl18_46 | ~ spl18_87)), inference(subsumption_resolution, [], [f1704, f190])).
fof(f1704, plain, ((e1 = e2) | (~ spl18_37 | ~ spl18_46 | ~ spl18_87)), inference(forward_demodulation, [], [f1701, f555])).
fof(f1701, plain, ((e2 = op(e1, e0)) | (~ spl18_37 | ~ spl18_87)), inference(backward_demodulation, [], [f857, f517])).
fof(f857, plain, ((e2 = op(e1, op(e1, e2))) | ~ spl18_87), inference(avatar_component_clause, [], [f855])).
fof(f1687, plain, (~ spl18_33 | ~ spl18_49), inference(avatar_split_clause, [], [f1681, f566, f498])).
fof(f1681, plain, (~ (e0 = op(e1, e3)) | ~ spl18_49), inference(backward_demodulation, [], [f157, f568])).
fof(f1665, plain, (~ spl18_44 | ~ spl18_22 | spl18_104), inference(avatar_split_clause, [], [f1522, f948, f451, f544])).
fof(f1522, plain, (~ (e3 = op(e1, e1)) | (~ spl18_22 | spl18_104)), inference(backward_demodulation, [], [f950, f453])).
fof(f1660, plain, (~ spl18_3 | ~ spl18_64 | spl18_114), inference(avatar_split_clause, [], [f1597, f993, f629, f370])).
fof(f1597, plain, (~ (e2 = op(e3, e3)) | (~ spl18_64 | spl18_114)), inference(backward_demodulation, [], [f995, f631])).
fof(f1647, plain, (~ spl18_39 | ~ spl18_40), inference(avatar_contradiction_clause, [], [f1646])).
fof(f1646, plain, ($false | (~ spl18_39 | ~ spl18_40)), inference(subsumption_resolution, [], [f1645, f192])).
fof(f1645, plain, ((e2 = e3) | (~ spl18_39 | ~ spl18_40)), inference(backward_demodulation, [], [f529, f525])).
fof(f1644, plain, (spl18_35 | ~ spl18_40 | ~ spl18_87), inference(avatar_contradiction_clause, [], [f1643])).
fof(f1643, plain, ($false | (spl18_35 | ~ spl18_40 | ~ spl18_87)), inference(subsumption_resolution, [], [f1641, f507])).
fof(f507, plain, (~ (e2 = op(e1, e3)) | spl18_35), inference(avatar_component_clause, [], [f506])).
fof(f1641, plain, ((e2 = op(e1, e3)) | (~ spl18_40 | ~ spl18_87)), inference(backward_demodulation, [], [f857, f529])).
fof(f1615, plain, (~ spl18_55 | ~ spl18_59), inference(avatar_split_clause, [], [f1611, f608, f591])).
fof(f1611, plain, (~ (e2 = op(e0, e2)) | ~ spl18_59), inference(backward_demodulation, [], [f166, f610])).
fof(f1602, plain, (~ spl18_56 | ~ spl18_64), inference(avatar_split_clause, [], [f1593, f629, f595])).
fof(f1593, plain, (~ (e3 = op(e0, e2)) | ~ spl18_64), inference(backward_demodulation, [], [f164, f631])).
fof(f1588, plain, (~ spl18_30 | ~ spl18_25 | spl18_83), inference(avatar_split_clause, [], [f1587, f833, f464, f485])).
fof(f1587, plain, (~ (e1 = op(e2, e0)) | (~ spl18_25 | spl18_83)), inference(forward_demodulation, [], [f834, f466])).
fof(f1563, plain, (~ spl18_37 | ~ spl18_3 | ~ spl18_22 | spl18_117), inference(avatar_split_clause, [], [f1550, f1008, f451, f370, f515])).
fof(f1550, plain, (~ (e0 = op(e1, e2)) | (~ spl18_3 | ~ spl18_22 | spl18_117)), inference(forward_demodulation, [], [f1545, f453])).
fof(f1545, plain, (~ (e0 = op(op(e2, e2), e2)) | (~ spl18_3 | spl18_117)), inference(backward_demodulation, [], [f1010, f372])).
fof(f1549, plain, (~ spl18_3 | ~ spl18_22 | spl18_110), inference(avatar_contradiction_clause, [], [f1548])).
fof(f1548, plain, ($false | (~ spl18_3 | ~ spl18_22 | spl18_110)), inference(subsumption_resolution, [], [f1544, f453])).
fof(f1544, plain, (~ (e1 = op(e2, e2)) | (~ spl18_3 | spl18_110)), inference(backward_demodulation, [], [f977, f372])).
fof(f1529, plain, (~ spl18_22 | ~ spl18_25 | ~ spl18_82), inference(avatar_contradiction_clause, [], [f1528])).
fof(f1528, plain, ($false | (~ spl18_22 | ~ spl18_25 | ~ spl18_82)), inference(subsumption_resolution, [], [f1527, f188])).
fof(f1527, plain, ((e0 = e2) | (~ spl18_22 | ~ spl18_25 | ~ spl18_82)), inference(forward_demodulation, [], [f1521, f466])).
fof(f1521, plain, ((e2 = op(e2, e1)) | (~ spl18_22 | ~ spl18_82)), inference(backward_demodulation, [], [f830, f453])).
fof(f1515, plain, (~ spl18_30 | ~ spl18_31), inference(avatar_contradiction_clause, [], [f1514])).
fof(f1514, plain, ($false | (~ spl18_30 | ~ spl18_31)), inference(subsumption_resolution, [], [f1513, f190])).
fof(f1513, plain, ((e1 = e2) | (~ spl18_30 | ~ spl18_31)), inference(backward_demodulation, [], [f491, f487])).
fof(f1512, plain, (~ spl18_34 | ~ spl18_43 | ~ spl18_86), inference(avatar_contradiction_clause, [], [f1511])).
fof(f1511, plain, ($false | (~ spl18_34 | ~ spl18_43 | ~ spl18_86)), inference(subsumption_resolution, [], [f1510, f192])).
fof(f1510, plain, ((e2 = e3) | (~ spl18_34 | ~ spl18_43 | ~ spl18_86)), inference(forward_demodulation, [], [f1509, f542])).
fof(f1509, plain, ((e3 = op(e1, e1)) | (~ spl18_34 | ~ spl18_86)), inference(backward_demodulation, [], [f852, f504])).
fof(f1505, plain, (~ spl18_37 | ~ spl18_38), inference(avatar_contradiction_clause, [], [f1504])).
fof(f1504, plain, ($false | (~ spl18_37 | ~ spl18_38)), inference(subsumption_resolution, [], [f1503, f187])).
fof(f1503, plain, ((e0 = e1) | (~ spl18_37 | ~ spl18_38)), inference(backward_demodulation, [], [f521, f517])).
fof(f1473, plain, (~ spl18_49 | spl18_64 | ~ spl18_91), inference(avatar_contradiction_clause, [], [f1472])).
fof(f1472, plain, ($false | (~ spl18_49 | spl18_64 | ~ spl18_91)), inference(subsumption_resolution, [], [f1471, f630])).
fof(f630, plain, (~ (op(e0, e0) = e3) | spl18_64), inference(avatar_component_clause, [], [f629])).
fof(f1471, plain, ((op(e0, e0) = e3) | (~ spl18_49 | ~ spl18_91)), inference(forward_demodulation, [], [f879, f568])).
fof(f1469, plain, (~ spl18_49 | ~ spl18_60 | ~ spl18_93), inference(avatar_contradiction_clause, [], [f1468])).
fof(f1468, plain, ($false | (~ spl18_49 | ~ spl18_60 | ~ spl18_93)), inference(subsumption_resolution, [], [f1467, f187])).
fof(f1467, plain, ((e0 = e1) | (~ spl18_49 | ~ spl18_60 | ~ spl18_93)), inference(forward_demodulation, [], [f1466, f568])).
fof(f1466, plain, ((e1 = op(e0, e3)) | (~ spl18_60 | ~ spl18_93)), inference(forward_demodulation, [], [f889, f614])).
fof(f1455, plain, (~ spl18_44 | ~ spl18_48), inference(avatar_split_clause, [], [f1454, f561, f544])).
fof(f1454, plain, (~ (e3 = op(e1, e1)) | ~ spl18_48), inference(forward_demodulation, [], [f169, f563])).
fof(f1451, plain, (spl18_43 | ~ spl18_38 | ~ spl18_87), inference(avatar_split_clause, [], [f1382, f855, f519, f540])).
fof(f1382, plain, ((e2 = op(e1, e1)) | (~ spl18_38 | ~ spl18_87)), inference(forward_demodulation, [], [f857, f521])).
fof(f1418, plain, (~ spl18_5 | ~ spl18_8), inference(avatar_contradiction_clause, [], [f1417])).
fof(f1417, plain, ($false | (~ spl18_5 | ~ spl18_8)), inference(subsumption_resolution, [], [f1415, f189])).
fof(f1415, plain, ((e0 = e3) | (~ spl18_5 | ~ spl18_8)), inference(backward_demodulation, [], [f393, f381])).
fof(f1411, plain, (~ spl18_29 | ~ spl18_31), inference(avatar_contradiction_clause, [], [f1410])).
fof(f1410, plain, ($false | (~ spl18_29 | ~ spl18_31)), inference(subsumption_resolution, [], [f1409, f188])).
fof(f1409, plain, ((e0 = e2) | (~ spl18_29 | ~ spl18_31)), inference(backward_demodulation, [], [f491, f483])).
fof(f1405, plain, (~ spl18_49 | ~ spl18_50), inference(avatar_contradiction_clause, [], [f1404])).
fof(f1404, plain, ($false | (~ spl18_49 | ~ spl18_50)), inference(subsumption_resolution, [], [f1403, f187])).
fof(f1403, plain, ((e0 = e1) | (~ spl18_49 | ~ spl18_50)), inference(backward_demodulation, [], [f572, f568])).
fof(f1399, plain, (~ spl18_6 | ~ spl18_24 | ~ spl18_63 | spl18_119), inference(avatar_split_clause, [], [f1398, f1018, f625, f459, f383])).
fof(f1398, plain, (~ (e1 = op(e3, e2)) | (~ spl18_24 | ~ spl18_63 | spl18_119)), inference(forward_demodulation, [], [f1393, f461])).
fof(f1393, plain, (~ (e1 = op(op(e2, e2), e2)) | (~ spl18_63 | spl18_119)), inference(backward_demodulation, [], [f1020, f627])).
fof(f1389, plain, (~ spl18_35 | ~ spl18_38 | ~ spl18_86), inference(avatar_contradiction_clause, [], [f1388])).
fof(f1388, plain, ($false | (~ spl18_35 | ~ spl18_38 | ~ spl18_86)), inference(subsumption_resolution, [], [f1387, f191])).
fof(f1387, plain, ((e1 = e3) | (~ spl18_35 | ~ spl18_38 | ~ spl18_86)), inference(forward_demodulation, [], [f1386, f521])).
fof(f1386, plain, ((e3 = op(e1, e2)) | (~ spl18_35 | ~ spl18_86)), inference(forward_demodulation, [], [f852, f508])).
fof(f1385, plain, (~ spl18_38 | ~ spl18_41 | ~ spl18_87), inference(avatar_contradiction_clause, [], [f1384])).
fof(f1384, plain, ($false | (~ spl18_38 | ~ spl18_41 | ~ spl18_87)), inference(subsumption_resolution, [], [f1383, f188])).
fof(f1383, plain, ((e0 = e2) | (~ spl18_38 | ~ spl18_41 | ~ spl18_87)), inference(forward_demodulation, [], [f1382, f534])).
fof(f1381, plain, (~ spl18_41 | spl18_46 | ~ spl18_88), inference(avatar_contradiction_clause, [], [f1380])).
fof(f1380, plain, ($false | (~ spl18_41 | spl18_46 | ~ spl18_88)), inference(subsumption_resolution, [], [f1379, f554])).
fof(f554, plain, (~ (e1 = op(e1, e0)) | spl18_46), inference(avatar_component_clause, [], [f553])).
fof(f1379, plain, ((e1 = op(e1, e0)) | (~ spl18_41 | ~ spl18_88)), inference(forward_demodulation, [], [f862, f534])).
fof(f1371, plain, (~ spl18_54 | ~ spl18_38), inference(avatar_split_clause, [], [f1370, f519, f587])).
fof(f1370, plain, (~ (e1 = op(e0, e2)) | ~ spl18_38), inference(forward_demodulation, [], [f151, f521])).
fof(f1365, plain, (~ spl18_54 | ~ spl18_50), inference(avatar_split_clause, [], [f1364, f570, f587])).
fof(f1364, plain, (~ (e1 = op(e0, e2)) | ~ spl18_50), inference(forward_demodulation, [], [f168, f572])).
fof(f1353, plain, (spl18_19 | ~ spl18_24 | ~ spl18_82), inference(avatar_contradiction_clause, [], [f1352])).
fof(f1352, plain, ($false | (spl18_19 | ~ spl18_24 | ~ spl18_82)), inference(subsumption_resolution, [], [f1316, f439])).
fof(f439, plain, (~ (e2 = op(e2, e3)) | spl18_19), inference(avatar_component_clause, [], [f438])).
fof(f1316, plain, ((e2 = op(e2, e3)) | (~ spl18_24 | ~ spl18_82)), inference(backward_demodulation, [], [f830, f461])).
fof(f1345, plain, (~ spl18_6 | ~ spl18_38), inference(avatar_split_clause, [], [f1344, f519, f383])).
fof(f1344, plain, (~ (e1 = op(e3, e2)) | ~ spl18_38), inference(forward_demodulation, [], [f155, f521])).
fof(f1338, plain, (~ spl18_4 | ~ spl18_24 | spl18_104), inference(avatar_split_clause, [], [f1317, f948, f459, f374])).
fof(f1317, plain, (~ (e3 = op(e3, e3)) | (~ spl18_24 | spl18_104)), inference(backward_demodulation, [], [f950, f461])).
fof(f1334, plain, (~ spl18_13 | ~ spl18_14), inference(avatar_contradiction_clause, [], [f1333])).
fof(f1333, plain, ($false | (~ spl18_13 | ~ spl18_14)), inference(subsumption_resolution, [], [f1332, f187])).
fof(f1332, plain, ((e0 = e1) | (~ spl18_13 | ~ spl18_14)), inference(backward_demodulation, [], [f419, f415])).
fof(f1300, plain, (~ spl18_53 | ~ spl18_55), inference(avatar_contradiction_clause, [], [f1299])).
fof(f1299, plain, ($false | (~ spl18_53 | ~ spl18_55)), inference(subsumption_resolution, [], [f1298, f188])).
fof(f1298, plain, ((e0 = e2) | (~ spl18_53 | ~ spl18_55)), inference(forward_demodulation, [], [f593, f585])).
fof(f1289, plain, (~ spl18_57 | ~ spl18_61), inference(avatar_split_clause, [], [f1279, f617, f600])).
fof(f1279, plain, (~ (e0 = op(e0, e1)) | ~ spl18_61), inference(backward_demodulation, [], [f163, f619])).
fof(f1287, plain, (~ spl18_45 | ~ spl18_61), inference(avatar_split_clause, [], [f1277, f617, f549])).
fof(f1277, plain, (~ (e0 = op(e1, e0)) | ~ spl18_61), inference(backward_demodulation, [], [f139, f619])).
fof(f1270, plain, (~ spl18_59 | ~ spl18_11), inference(avatar_split_clause, [], [f1269, f404, f608])).
fof(f1269, plain, (~ (e2 = op(e0, e1)) | ~ spl18_11), inference(forward_demodulation, [], [f147, f406])).
fof(f1255, plain, (~ spl18_47 | ~ spl18_35), inference(avatar_split_clause, [], [f1254, f506, f557])).
fof(f1254, plain, (~ (e2 = op(e1, e0)) | ~ spl18_35), inference(forward_demodulation, [], [f171, f508])).
fof(f1244, plain, (~ spl18_30 | ~ spl18_14), inference(avatar_split_clause, [], [f1243, f417, f485])).
fof(f1243, plain, (~ (e1 = op(e2, e0)) | ~ spl18_14), inference(forward_demodulation, [], [f144, f419])).
fof(f1240, plain, (~ spl18_27 | ~ spl18_11), inference(avatar_split_clause, [], [f1239, f404, f472])).
fof(f1239, plain, (~ (e2 = op(e2, e1)) | ~ spl18_11), inference(forward_demodulation, [], [f150, f406])).
fof(f1233, plain, (~ spl18_17 | ~ spl18_20), inference(avatar_contradiction_clause, [], [f1232])).
fof(f1232, plain, ($false | (~ spl18_17 | ~ spl18_20)), inference(subsumption_resolution, [], [f1231, f189])).
fof(f1231, plain, ((e0 = e3) | (~ spl18_17 | ~ spl18_20)), inference(forward_demodulation, [], [f444, f432])).
fof(f1228, plain, (~ spl18_6 | ~ spl18_14), inference(avatar_contradiction_clause, [], [f1227])).
fof(f1227, plain, ($false | (~ spl18_6 | ~ spl18_14)), inference(subsumption_resolution, [], [f1226, f419])).
fof(f1226, plain, (~ (e1 = op(e3, e0)) | ~ spl18_6), inference(forward_demodulation, [], [f182, f385])).
fof(f1218, plain, (~ spl18_6 | ~ spl18_8), inference(avatar_contradiction_clause, [], [f1217])).
fof(f1217, plain, ($false | (~ spl18_6 | ~ spl18_8)), inference(subsumption_resolution, [], [f1216, f191])).
fof(f1216, plain, ((e1 = e3) | (~ spl18_6 | ~ spl18_8)), inference(forward_demodulation, [], [f393, f385])).
fof(f1203, plain, (~ spl18_5 | ~ spl18_6), inference(avatar_contradiction_clause, [], [f1202])).
fof(f1202, plain, ($false | (~ spl18_5 | ~ spl18_6)), inference(subsumption_resolution, [], [f1201, f187])).
fof(f1201, plain, ((e0 = e1) | (~ spl18_5 | ~ spl18_6)), inference(backward_demodulation, [], [f385, f381])).
fof(f1200, plain, (~ spl18_6 | ~ spl18_7), inference(avatar_contradiction_clause, [], [f1199])).
fof(f1199, plain, ($false | (~ spl18_6 | ~ spl18_7)), inference(subsumption_resolution, [], [f1198, f190])).
fof(f1198, plain, ((e1 = e2) | (~ spl18_6 | ~ spl18_7)), inference(backward_demodulation, [], [f389, f385])).
fof(f1197, plain, (~ spl18_7 | ~ spl18_8), inference(avatar_contradiction_clause, [], [f1196])).
fof(f1196, plain, ($false | (~ spl18_7 | ~ spl18_8)), inference(subsumption_resolution, [], [f1195, f192])).
fof(f1195, plain, ((e2 = e3) | (~ spl18_7 | ~ spl18_8)), inference(backward_demodulation, [], [f393, f389])).
fof(f1194, plain, (~ spl18_4 | ~ spl18_8), inference(avatar_split_clause, [], [f1193, f391, f374])).
fof(f1193, plain, (~ (e3 = op(e3, e3)) | ~ spl18_8), inference(backward_demodulation, [], [f186, f393])).
fof(f1192, plain, (~ spl18_10 | ~ spl18_11), inference(avatar_contradiction_clause, [], [f1191])).
fof(f1191, plain, ($false | (~ spl18_10 | ~ spl18_11)), inference(subsumption_resolution, [], [f1190, f190])).
fof(f1190, plain, ((e1 = e2) | (~ spl18_10 | ~ spl18_11)), inference(backward_demodulation, [], [f406, f402])).
fof(f1189, plain, (~ spl18_11 | ~ spl18_12), inference(avatar_contradiction_clause, [], [f1188])).
fof(f1188, plain, ($false | (~ spl18_11 | ~ spl18_12)), inference(subsumption_resolution, [], [f1187, f192])).
fof(f1187, plain, ((e2 = e3) | (~ spl18_11 | ~ spl18_12)), inference(backward_demodulation, [], [f410, f406])).
fof(f1182, plain, (~ spl18_14 | ~ spl18_15), inference(avatar_contradiction_clause, [], [f1181])).
fof(f1181, plain, ($false | (~ spl18_14 | ~ spl18_15)), inference(subsumption_resolution, [], [f1180, f190])).
fof(f1180, plain, ((e1 = e2) | (~ spl18_14 | ~ spl18_15)), inference(backward_demodulation, [], [f423, f419])).
fof(f1177, plain, (~ spl18_11 | ~ spl18_15), inference(avatar_split_clause, [], [f1174, f421, f404])).
fof(f1174, plain, (~ (e2 = op(e3, e1)) | ~ spl18_15), inference(backward_demodulation, [], [f181, f423])).
fof(f1173, plain, (~ spl18_17 | ~ spl18_18), inference(avatar_contradiction_clause, [], [f1172])).
fof(f1172, plain, ($false | (~ spl18_17 | ~ spl18_18)), inference(subsumption_resolution, [], [f1171, f187])).
fof(f1171, plain, ((e0 = e1) | (~ spl18_17 | ~ spl18_18)), inference(backward_demodulation, [], [f436, f432])).
fof(f1170, plain, (~ spl18_18 | ~ spl18_19), inference(avatar_contradiction_clause, [], [f1169])).
fof(f1169, plain, ($false | (~ spl18_18 | ~ spl18_19)), inference(subsumption_resolution, [], [f1168, f190])).
fof(f1168, plain, ((e1 = e2) | (~ spl18_18 | ~ spl18_19)), inference(backward_demodulation, [], [f440, f436])).
fof(f1167, plain, (~ spl18_19 | ~ spl18_20), inference(avatar_contradiction_clause, [], [f1166])).
fof(f1166, plain, ($false | (~ spl18_19 | ~ spl18_20)), inference(subsumption_resolution, [], [f1165, f192])).
fof(f1165, plain, ((e2 = e3) | (~ spl18_19 | ~ spl18_20)), inference(backward_demodulation, [], [f444, f440])).
fof(f1147, plain, (spl18_30 | ~ spl18_25 | ~ spl18_83), inference(avatar_split_clause, [], [f1143, f833, f464, f485])).
fof(f1143, plain, ((e1 = op(e2, e0)) | (~ spl18_25 | ~ spl18_83)), inference(backward_demodulation, [], [f835, f466])).
fof(f1139, plain, (spl18_17 | ~ spl18_32 | ~ spl18_84), inference(avatar_split_clause, [], [f1134, f838, f493, f430])).
fof(f1134, plain, ((e0 = op(e2, e3)) | (~ spl18_32 | ~ spl18_84)), inference(backward_demodulation, [], [f840, f495])).
fof(f1129, plain, (~ spl18_34 | ~ spl18_35), inference(avatar_contradiction_clause, [], [f1128])).
fof(f1128, plain, ($false | (~ spl18_34 | ~ spl18_35)), inference(subsumption_resolution, [], [f1127, f190])).
fof(f1127, plain, ((e1 = e2) | (~ spl18_34 | ~ spl18_35)), inference(backward_demodulation, [], [f508, f504])).
fof(f1126, plain, (~ spl18_35 | ~ spl18_36), inference(avatar_contradiction_clause, [], [f1125])).
fof(f1125, plain, ($false | (~ spl18_35 | ~ spl18_36)), inference(subsumption_resolution, [], [f1124, f192])).
fof(f1124, plain, ((e2 = e3) | (~ spl18_35 | ~ spl18_36)), inference(backward_demodulation, [], [f512, f508])).
fof(f1111, plain, (~ spl18_15 | ~ spl18_41 | ~ spl18_64 | spl18_99), inference(avatar_split_clause, [], [f1110, f926, f629, f532, f421])).
fof(f1110, plain, (~ (e2 = op(e3, e0)) | (~ spl18_41 | ~ spl18_64 | spl18_99)), inference(forward_demodulation, [], [f1104, f631])).
fof(f1104, plain, (~ (e2 = op(op(e0, e0), e0)) | (~ spl18_41 | spl18_99)), inference(backward_demodulation, [], [f928, f534])).
fof(f1085, plain, (~ spl18_50 | ~ spl18_51), inference(avatar_contradiction_clause, [], [f1084])).
fof(f1084, plain, ($false | (~ spl18_50 | ~ spl18_51)), inference(subsumption_resolution, [], [f1083, f190])).
fof(f1083, plain, ((e1 = e2) | (~ spl18_50 | ~ spl18_51)), inference(backward_demodulation, [], [f576, f572])).
fof(f1080, plain, (~ spl18_35 | ~ spl18_51), inference(avatar_split_clause, [], [f1077, f574, f506])).
fof(f1077, plain, (~ (e2 = op(e1, e3)) | ~ spl18_51), inference(backward_demodulation, [], [f157, f576])).
fof(f1074, plain, (~ spl18_21 | ~ spl18_53), inference(avatar_split_clause, [], [f1070, f583, f447])).
fof(f1070, plain, (~ (e0 = op(e2, e2)) | ~ spl18_53), inference(backward_demodulation, [], [f152, f585])).
fof(f1054, plain, (~ spl18_32 | ~ spl18_64), inference(avatar_split_clause, [], [f1046, f629, f493])).
fof(f1046, plain, (~ (e3 = op(e2, e0)) | ~ spl18_64), inference(backward_demodulation, [], [f140, f631])).
fof(f1044, plain, (~ spl18_24 | ~ spl18_103 | ~ spl18_116), inference(avatar_split_clause, [], [f344, f1003, f944, f459])).
fof(f344, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | ~ (e1 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | ~ (e1 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e0 = op(op(e2, e2), op(e2, e2))) & (e1 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) & (e3 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax29)).
fof(f1043, plain, (~ spl18_24 | ~ spl18_115 | ~ spl18_106), inference(avatar_split_clause, [], [f343, f957, f998, f459])).
fof(f343, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | ~ (e0 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | ~ (e0 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e1 = op(op(e2, e2), op(e2, e2))) & (e0 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) & (e3 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax28)).
fof(f1042, plain, (~ spl18_44 | ~ spl18_99 | ~ spl18_122), inference(avatar_split_clause, [], [f342, f1033, f926, f544])).
fof(f342, plain, (~ (e0 = op(op(e1, e1), op(e1, e1))) | ~ (e2 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e3 = op(e1, e1))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e0 = op(op(e1, e1), op(e1, e1))) | ~ (e2 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e3 = op(e1, e1))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e0 = op(op(e1, e1), op(e1, e1))) & (e2 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) & (e3 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax27)).
fof(f1041, plain, (~ spl18_44 | ~ spl18_121 | ~ spl18_102), inference(avatar_split_clause, [], [f341, f939, f1028, f544])).
fof(f341, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | ~ (e0 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e3 = op(e1, e1))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | ~ (e0 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e3 = op(e1, e1))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(op(e1, e1), op(e1, e1))) & (e0 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) & (e3 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax26)).
fof(f1040, plain, (~ spl18_64 | ~ spl18_111 | ~ spl18_120), inference(avatar_split_clause, [], [f340, f1023, f980, f629])).
fof(f340, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (e2 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e3)), inference(cnf_transformation, [], [f49])).
fof(f49, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (e2 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e3)), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e1 = op(op(e0, e0), op(e0, e0))) & (e2 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) & (op(e0, e0) = e3)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax25)).
fof(f1039, plain, (~ spl18_64 | ~ spl18_119 | ~ spl18_114), inference(avatar_split_clause, [], [f339, f993, f1018, f629])).
fof(f339, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | ~ (e1 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e3)), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | ~ (e1 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e3)), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e2 = op(op(e0, e0), op(e0, e0))) & (e1 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) & (op(e0, e0) = e3)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax24)).
fof(f1038, plain, (~ spl18_3 | ~ spl18_107 | ~ spl18_118), inference(avatar_split_clause, [], [f338, f1013, f962, f370])).
fof(f338, plain, (~ (e0 = op(op(e3, e3), op(e3, e3))) | ~ (e1 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e2 = op(e3, e3))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e0 = op(op(e3, e3), op(e3, e3))) | ~ (e1 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e2 = op(e3, e3))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(op(e3, e3), op(e3, e3))) & (e1 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) & (e2 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax23)).
fof(f1037, plain, (~ spl18_3 | ~ spl18_117 | ~ spl18_110), inference(avatar_split_clause, [], [f337, f975, f1008, f370])).
fof(f337, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | ~ (e0 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e2 = op(e3, e3))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | ~ (e0 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e2 = op(e3, e3))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(op(e3, e3), op(e3, e3))) & (e0 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) & (e2 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax22)).
fof(f1036, plain, (~ spl18_43 | ~ spl18_101 | ~ spl18_122), inference(avatar_split_clause, [], [f336, f1033, f935, f540])).
fof(f336, plain, (~ (e0 = op(op(e1, e1), op(e1, e1))) | ~ (e3 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e2 = op(e1, e1))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(op(e1, e1), op(e1, e1))) | ~ (e3 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e2 = op(e1, e1))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(op(e1, e1), op(e1, e1))) & (e3 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) & (e2 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax21)).
fof(f1031, plain, (~ spl18_43 | ~ spl18_121 | ~ spl18_100), inference(avatar_split_clause, [], [f335, f930, f1028, f540])).
fof(f335, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | ~ (e0 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e2 = op(e1, e1))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | ~ (e0 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e2 = op(e1, e1))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e3 = op(op(e1, e1), op(e1, e1))) & (e0 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) & (e2 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax20)).
fof(f1026, plain, (~ spl18_63 | ~ spl18_113 | ~ spl18_120), inference(avatar_split_clause, [], [f334, f1023, f989, f625])).
fof(f334, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (e3 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e2)), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | ~ (e3 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e2)), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(op(e0, e0), op(e0, e0))) & (e3 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) & (op(e0, e0) = e2)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax19)).
fof(f1021, plain, (~ spl18_63 | ~ spl18_119 | ~ spl18_112), inference(avatar_split_clause, [], [f333, f984, f1018, f625])).
fof(f333, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (e1 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e2)), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (e1 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e2)), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(op(e0, e0), op(e0, e0))) & (e1 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) & (op(e0, e0) = e2)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax18)).
fof(f1016, plain, (~ spl18_2 | ~ spl18_109 | ~ spl18_118), inference(avatar_split_clause, [], [f332, f1013, f971, f366])).
fof(f332, plain, (~ (e0 = op(op(e3, e3), op(e3, e3))) | ~ (e2 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e1 = op(e3, e3))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e0 = op(op(e3, e3), op(e3, e3))) | ~ (e2 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e1 = op(e3, e3))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(op(e3, e3), op(e3, e3))) & (e2 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) & (e1 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax17)).
fof(f1011, plain, (~ spl18_2 | ~ spl18_117 | ~ spl18_108), inference(avatar_split_clause, [], [f331, f966, f1008, f366])).
fof(f331, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | ~ (e0 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e1 = op(e3, e3))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | ~ (e0 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e1 = op(e3, e3))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e2 = op(op(e3, e3), op(e3, e3))) & (e0 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) & (e1 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax16)).
fof(f1006, plain, (~ spl18_22 | ~ spl18_105 | ~ spl18_116), inference(avatar_split_clause, [], [f330, f1003, f953, f451])).
fof(f330, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | ~ (e3 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | ~ (e3 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(op(e2, e2), op(e2, e2))) & (e3 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) & (e1 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax15)).
fof(f1001, plain, (~ spl18_22 | ~ spl18_115 | ~ spl18_104), inference(avatar_split_clause, [], [f329, f948, f998, f451])).
fof(f329, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | ~ (e0 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | ~ (e0 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(op(e2, e2), op(e2, e2))) & (e0 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) & (e1 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax14)).
fof(f996, plain, (~ spl18_62 | ~ spl18_113 | ~ spl18_114), inference(avatar_split_clause, [], [f328, f993, f989, f621])).
fof(f328, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | ~ (e3 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e1)), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | ~ (e3 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e1)), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e2 = op(op(e0, e0), op(e0, e0))) & (e3 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) & (op(e0, e0) = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax13)).
fof(f987, plain, (~ spl18_62 | ~ spl18_111 | ~ spl18_112), inference(avatar_split_clause, [], [f327, f984, f980, f621])).
fof(f327, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (e2 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e1)), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | ~ (e2 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) | ~ (op(e0, e0) = e1)), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(op(e0, e0), op(e0, e0))) & (e2 = op(op(op(e0, e0), op(e0, e0)), op(e0, e0))) & (op(e0, e0) = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax12)).
fof(f978, plain, (~ spl18_1 | ~ spl18_109 | ~ spl18_110), inference(avatar_split_clause, [], [f326, f975, f971, f362])).
fof(f326, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | ~ (e2 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e0 = op(e3, e3))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | ~ (e2 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e0 = op(e3, e3))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e1 = op(op(e3, e3), op(e3, e3))) & (e2 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) & (e0 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax11)).
fof(f969, plain, (~ spl18_1 | ~ spl18_107 | ~ spl18_108), inference(avatar_split_clause, [], [f325, f966, f962, f362])).
fof(f325, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | ~ (e1 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e0 = op(e3, e3))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | ~ (e1 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) | ~ (e0 = op(e3, e3))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(op(e3, e3), op(e3, e3))) & (e1 = op(op(op(e3, e3), op(e3, e3)), op(e3, e3))) & (e0 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax10)).
fof(f960, plain, (~ spl18_21 | ~ spl18_105 | ~ spl18_106), inference(avatar_split_clause, [], [f324, f957, f953, f447])).
fof(f324, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | ~ (e3 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e0 = op(e2, e2))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | ~ (e3 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e0 = op(e2, e2))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e1 = op(op(e2, e2), op(e2, e2))) & (e3 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) & (e0 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax9)).
fof(f951, plain, (~ spl18_21 | ~ spl18_103 | ~ spl18_104), inference(avatar_split_clause, [], [f323, f948, f944, f447])).
fof(f323, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | ~ (e1 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e0 = op(e2, e2))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e3 = op(op(e2, e2), op(e2, e2))) | ~ (e1 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) | ~ (e0 = op(e2, e2))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(op(e2, e2), op(e2, e2))) & (e1 = op(op(op(e2, e2), op(e2, e2)), op(e2, e2))) & (e0 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax8)).
fof(f942, plain, (~ spl18_41 | ~ spl18_101 | ~ spl18_102), inference(avatar_split_clause, [], [f322, f939, f935, f532])).
fof(f322, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | ~ (e3 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e0 = op(e1, e1))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | ~ (e3 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e0 = op(e1, e1))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e2 = op(op(e1, e1), op(e1, e1))) & (e3 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) & (e0 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax7)).
fof(f933, plain, (~ spl18_41 | ~ spl18_99 | ~ spl18_100), inference(avatar_split_clause, [], [f321, f930, f926, f532])).
fof(f321, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | ~ (e2 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e0 = op(e1, e1))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | ~ (e2 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) | ~ (e0 = op(e1, e1))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e3 = op(op(e1, e1), op(e1, e1))) & (e2 = op(op(op(e1, e1), op(e1, e1)), op(e1, e1))) & (e0 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax6)).
fof(f924, plain, (spl18_90 | spl18_85 | spl18_80 | spl18_98), inference(avatar_split_clause, [], [f307, f921, f816, f843, f870])).
fof(f870, plain, (spl18_90 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl18_90])])).
fof(f843, plain, (spl18_85 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl18_85])])).
fof(f816, plain, (spl18_80 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl18_80])])).
fof(f307, plain, ((e0 = op(e3, op(e3, e0))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f72])).
fof(f72, plain, ((((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3) & ((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2))) & ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1))) & ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3)) & (e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0)))) | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0)))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0)))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, (((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> ((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, (((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> ((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, (((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> ((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, (((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> ((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, (((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> ((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, (((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> ((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, (((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> ((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, (((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> ((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, (((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> ((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, (((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> ((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, (((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> ((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, (((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> ((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, (((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> ((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, (((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> ((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, (((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> ((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f5, plain, ((((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2))) & ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1))) & ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3)) & (e3 = op(e3, op(e3, e3))) & (e2 = op(e3, op(e3, e2))) & (e1 = op(e3, op(e3, e1))) & (e0 = op(e3, op(e3, e0)))) | (((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0)))) | (((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0)))) | (((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0)))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax5)).
fof(f919, plain, (spl18_90 | spl18_85 | spl18_80 | spl18_97), inference(avatar_split_clause, [], [f308, f916, f816, f843, f870])).
fof(f308, plain, ((e1 = op(e3, op(e3, e1))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f72])).
fof(f914, plain, (spl18_90 | spl18_85 | spl18_80 | spl18_96), inference(avatar_split_clause, [], [f309, f911, f816, f843, f870])).
fof(f309, plain, ((e2 = op(e3, op(e3, e2))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f72])).
fof(f909, plain, (spl18_90 | spl18_85 | spl18_80 | spl18_95), inference(avatar_split_clause, [], [f310, f906, f816, f843, f870])).
fof(f310, plain, ((e3 = op(e3, op(e3, e3))) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f72])).
fof(f904, plain, (spl18_90 | spl18_85 | spl18_80 | ~ spl18_64 | spl18_49), inference(avatar_split_clause, [], [f311, f566, f629, f816, f843, f870])).
fof(f311, plain, ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f72])).
fof(f903, plain, (spl18_90 | spl18_85 | spl18_80 | ~ spl18_44 | spl18_34), inference(avatar_split_clause, [], [f312, f502, f544, f816, f843, f870])).
fof(f312, plain, ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f72])).
fof(f902, plain, (spl18_90 | spl18_85 | spl18_80 | ~ spl18_24 | spl18_19), inference(avatar_split_clause, [], [f313, f438, f459, f816, f843, f870])).
fof(f313, plain, ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f72])).
fof(f901, plain, (spl18_79 | spl18_78 | spl18_77 | spl18_76 | spl18_75 | spl18_74 | spl18_73 | spl18_72 | spl18_71 | spl18_70 | spl18_69 | spl18_68 | spl18_67 | spl18_66 | spl18_65 | spl18_4), inference(avatar_split_clause, [], [f315, f374, f666, f676, f686, f696, f706, f716, f726, f736, f746, f756, f766, f776, f786, f796, f806])).
fof(f806, plain, (spl18_79 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl18_79])])).
fof(f796, plain, (spl18_78 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl18_78])])).
fof(f786, plain, (spl18_77 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl18_77])])).
fof(f776, plain, (spl18_76 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl18_76])])).
fof(f766, plain, (spl18_75 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl18_75])])).
fof(f756, plain, (spl18_74 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl18_74])])).
fof(f746, plain, (spl18_73 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl18_73])])).
fof(f736, plain, (spl18_72 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl18_72])])).
fof(f726, plain, (spl18_71 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl18_71])])).
fof(f716, plain, (spl18_70 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl18_70])])).
fof(f706, plain, (spl18_69 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl18_69])])).
fof(f696, plain, (spl18_68 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl18_68])])).
fof(f686, plain, (spl18_67 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl18_67])])).
fof(f676, plain, (spl18_66 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl18_66])])).
fof(f666, plain, (spl18_65 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl18_65])])).
fof(f315, plain, ((e3 = op(e3, e3)) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f72])).
fof(f899, plain, (spl18_79 | spl18_78 | spl18_77 | spl18_76 | spl18_75 | spl18_74 | spl18_73 | spl18_72 | spl18_71 | spl18_70 | spl18_69 | spl18_68 | spl18_67 | spl18_66 | spl18_65 | ~ spl18_16 | ~ spl18_1), inference(avatar_split_clause, [], [f317, f362, f425, f666, f676, f686, f696, f706, f716, f726, f736, f746, f756, f766, f776, f786, f796, f806])).
fof(f317, plain, (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f72])).
fof(f897, plain, (spl18_79 | spl18_78 | spl18_77 | spl18_76 | spl18_75 | spl18_74 | spl18_73 | spl18_72 | spl18_71 | spl18_70 | spl18_69 | spl18_68 | spl18_67 | spl18_66 | spl18_65 | ~ spl18_8 | ~ spl18_3), inference(avatar_split_clause, [], [f319, f370, f391, f666, f676, f686, f696, f706, f716, f726, f736, f746, f756, f766, f776, f786, f796, f806])).
fof(f319, plain, (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2)) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f72])).
fof(f896, plain, (spl18_79 | spl18_78 | spl18_77 | spl18_76 | spl18_75 | spl18_74 | spl18_73 | spl18_72 | spl18_71 | spl18_70 | spl18_69 | spl18_68 | spl18_67 | spl18_66 | spl18_65 | ~ spl18_4), inference(avatar_split_clause, [], [f345, f374, f666, f676, f686, f696, f706, f716, f726, f736, f746, f756, f766, f776, f786, f796, f806])).
fof(f345, plain, (~ (e3 = op(e3, e3)) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(duplicate_literal_removal, [], [f320])).
fof(f320, plain, (~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f72])).
fof(f895, plain, (~ spl18_90 | spl18_94), inference(avatar_split_clause, [], [f299, f892, f870])).
fof(f299, plain, ((e0 = op(e0, op(e0, e0))) | ~ sP0), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ((((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, op(e0, e3))) & (e2 = op(e0, op(e0, e2))) & (e1 = op(e0, op(e0, e1))) & (e0 = op(e0, op(e0, e0)))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f890, plain, (~ spl18_90 | spl18_93), inference(avatar_split_clause, [], [f300, f887, f870])).
fof(f300, plain, ((e1 = op(e0, op(e0, e1))) | ~ sP0), inference(cnf_transformation, [], [f90])).
fof(f885, plain, (~ spl18_90 | spl18_92), inference(avatar_split_clause, [], [f301, f882, f870])).
fof(f301, plain, ((e2 = op(e0, op(e0, e2))) | ~ sP0), inference(cnf_transformation, [], [f90])).
fof(f880, plain, (~ spl18_90 | spl18_91), inference(avatar_split_clause, [], [f302, f877, f870])).
fof(f302, plain, ((e3 = op(e0, op(e0, e3))) | ~ sP0), inference(cnf_transformation, [], [f90])).
fof(f875, plain, (~ spl18_90 | ~ spl18_41 | spl18_46), inference(avatar_split_clause, [], [f304, f553, f532, f870])).
fof(f304, plain, ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1)) | ~ sP0), inference(cnf_transformation, [], [f90])).
fof(f874, plain, (~ spl18_90 | ~ spl18_21 | spl18_31), inference(avatar_split_clause, [], [f305, f489, f447, f870])).
fof(f305, plain, ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2)) | ~ sP0), inference(cnf_transformation, [], [f90])).
fof(f873, plain, (~ spl18_90 | ~ spl18_1 | spl18_16), inference(avatar_split_clause, [], [f306, f425, f362, f870])).
fof(f306, plain, ((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3)) | ~ sP0), inference(cnf_transformation, [], [f90])).
fof(f868, plain, (~ spl18_85 | spl18_89), inference(avatar_split_clause, [], [f291, f865, f843])).
fof(f291, plain, ((e0 = op(e1, op(e1, e0))) | ~ sP1), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e3 = op(e1, op(e1, e3))) & (e2 = op(e1, op(e1, e2))) & (e1 = op(e1, op(e1, e1))) & (e0 = op(e1, op(e1, e0)))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f863, plain, (~ spl18_85 | spl18_88), inference(avatar_split_clause, [], [f292, f860, f843])).
fof(f292, plain, ((e1 = op(e1, op(e1, e1))) | ~ sP1), inference(cnf_transformation, [], [f89])).
fof(f858, plain, (~ spl18_85 | spl18_87), inference(avatar_split_clause, [], [f293, f855, f843])).
fof(f293, plain, ((e2 = op(e1, op(e1, e2))) | ~ sP1), inference(cnf_transformation, [], [f89])).
fof(f853, plain, (~ spl18_85 | spl18_86), inference(avatar_split_clause, [], [f294, f850, f843])).
fof(f294, plain, ((e3 = op(e1, op(e1, e3))) | ~ sP1), inference(cnf_transformation, [], [f89])).
fof(f848, plain, (~ spl18_85 | ~ spl18_62 | spl18_57), inference(avatar_split_clause, [], [f295, f600, f621, f843])).
fof(f295, plain, ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1) | ~ sP1), inference(cnf_transformation, [], [f89])).
fof(f847, plain, (~ spl18_85 | ~ spl18_22 | spl18_27), inference(avatar_split_clause, [], [f297, f472, f451, f843])).
fof(f297, plain, ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2)) | ~ sP1), inference(cnf_transformation, [], [f89])).
fof(f846, plain, (~ spl18_85 | ~ spl18_2 | spl18_12), inference(avatar_split_clause, [], [f298, f408, f366, f843])).
fof(f298, plain, ((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3)) | ~ sP1), inference(cnf_transformation, [], [f89])).
fof(f841, plain, (~ spl18_80 | spl18_84), inference(avatar_split_clause, [], [f283, f838, f816])).
fof(f283, plain, ((e0 = op(e2, op(e2, e0))) | ~ sP2), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e3 = op(e2, op(e2, e3))) & (e2 = op(e2, op(e2, e2))) & (e1 = op(e2, op(e2, e1))) & (e0 = op(e2, op(e2, e0)))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f836, plain, (~ spl18_80 | spl18_83), inference(avatar_split_clause, [], [f284, f833, f816])).
fof(f284, plain, ((e1 = op(e2, op(e2, e1))) | ~ sP2), inference(cnf_transformation, [], [f88])).
fof(f831, plain, (~ spl18_80 | spl18_82), inference(avatar_split_clause, [], [f285, f828, f816])).
fof(f285, plain, ((e2 = op(e2, op(e2, e2))) | ~ sP2), inference(cnf_transformation, [], [f88])).
fof(f826, plain, (~ spl18_80 | spl18_81), inference(avatar_split_clause, [], [f286, f823, f816])).
fof(f286, plain, ((e3 = op(e2, op(e2, e3))) | ~ sP2), inference(cnf_transformation, [], [f88])).
fof(f821, plain, (~ spl18_80 | ~ spl18_63 | spl18_53), inference(avatar_split_clause, [], [f287, f583, f625, f816])).
fof(f287, plain, ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2) | ~ sP2), inference(cnf_transformation, [], [f88])).
fof(f820, plain, (~ spl18_80 | ~ spl18_43 | spl18_38), inference(avatar_split_clause, [], [f288, f519, f540, f816])).
fof(f288, plain, ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1)) | ~ sP2), inference(cnf_transformation, [], [f88])).
fof(f819, plain, (~ spl18_80 | ~ spl18_3 | spl18_8), inference(avatar_split_clause, [], [f290, f391, f370, f816])).
fof(f290, plain, ((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3)) | ~ sP2), inference(cnf_transformation, [], [f88])).
fof(f814, plain, (~ spl18_79 | spl18_61), inference(avatar_split_clause, [], [f277, f617, f806])).
fof(f277, plain, ((e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f87])).
fof(f87, plain, (((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f812, plain, (~ spl18_79 | ~ spl18_61), inference(avatar_split_clause, [], [f346, f617, f806])).
fof(f346, plain, (~ (e0 = op(e0, e0)) | ~ sP3), inference(duplicate_literal_removal, [], [f279])).
fof(f279, plain, (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f87])).
fof(f804, plain, (~ spl18_78 | spl18_57), inference(avatar_split_clause, [], [f271, f600, f796])).
fof(f271, plain, ((e0 = op(e0, e1)) | ~ sP4), inference(cnf_transformation, [], [f86])).
fof(f86, plain, (((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e0 = op(e1, e0)) & (e0 = op(e0, e1))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f803, plain, (~ spl18_78 | spl18_45), inference(avatar_split_clause, [], [f272, f549, f796])).
fof(f272, plain, ((e0 = op(e1, e0)) | ~ sP4), inference(cnf_transformation, [], [f86])).
fof(f801, plain, (~ spl18_78 | ~ spl18_42), inference(avatar_split_clause, [], [f347, f536, f796])).
fof(f347, plain, (~ (e1 = op(e1, e1)) | ~ sP4), inference(duplicate_literal_removal, [], [f274])).
fof(f274, plain, (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ sP4), inference(cnf_transformation, [], [f86])).
fof(f800, plain, (~ spl18_78 | ~ spl18_38 | ~ spl18_43), inference(avatar_split_clause, [], [f275, f540, f519, f796])).
fof(f275, plain, (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2)) | ~ sP4), inference(cnf_transformation, [], [f86])).
fof(f799, plain, (~ spl18_78 | ~ spl18_34 | ~ spl18_44), inference(avatar_split_clause, [], [f276, f544, f502, f796])).
fof(f276, plain, (~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3)) | ~ sP4), inference(cnf_transformation, [], [f86])).
fof(f794, plain, (~ spl18_77 | spl18_53), inference(avatar_split_clause, [], [f265, f583, f786])).
fof(f265, plain, ((e0 = op(e0, e2)) | ~ sP5), inference(cnf_transformation, [], [f85])).
fof(f85, plain, (((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e0 = op(e2, e0)) & (e0 = op(e0, e2))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f793, plain, (~ spl18_77 | spl18_29), inference(avatar_split_clause, [], [f266, f481, f786])).
fof(f266, plain, ((e0 = op(e2, e0)) | ~ sP5), inference(cnf_transformation, [], [f85])).
fof(f792, plain, (~ spl18_77 | ~ spl18_31 | ~ spl18_21), inference(avatar_split_clause, [], [f267, f447, f489, f786])).
fof(f267, plain, (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP5), inference(cnf_transformation, [], [f85])).
fof(f791, plain, (~ spl18_77 | ~ spl18_27 | ~ spl18_22), inference(avatar_split_clause, [], [f268, f451, f472, f786])).
fof(f268, plain, (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1)) | ~ sP5), inference(cnf_transformation, [], [f85])).
fof(f790, plain, (~ spl18_77 | ~ spl18_23), inference(avatar_split_clause, [], [f348, f455, f786])).
fof(f348, plain, (~ (e2 = op(e2, e2)) | ~ sP5), inference(duplicate_literal_removal, [], [f269])).
fof(f269, plain, (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ sP5), inference(cnf_transformation, [], [f85])).
fof(f789, plain, (~ spl18_77 | ~ spl18_19 | ~ spl18_24), inference(avatar_split_clause, [], [f270, f459, f438, f786])).
fof(f270, plain, (~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3)) | ~ sP5), inference(cnf_transformation, [], [f85])).
fof(f784, plain, (~ spl18_76 | spl18_49), inference(avatar_split_clause, [], [f259, f566, f776])).
fof(f259, plain, ((e0 = op(e0, e3)) | ~ sP6), inference(cnf_transformation, [], [f84])).
fof(f84, plain, (((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e0 = op(e3, e0)) & (e0 = op(e0, e3))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f783, plain, (~ spl18_76 | spl18_13), inference(avatar_split_clause, [], [f260, f413, f776])).
fof(f260, plain, ((e0 = op(e3, e0)) | ~ sP6), inference(cnf_transformation, [], [f84])).
fof(f782, plain, (~ spl18_76 | ~ spl18_16 | ~ spl18_1), inference(avatar_split_clause, [], [f261, f362, f425, f776])).
fof(f261, plain, (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | ~ sP6), inference(cnf_transformation, [], [f84])).
fof(f781, plain, (~ spl18_76 | ~ spl18_12 | ~ spl18_2), inference(avatar_split_clause, [], [f262, f366, f408, f776])).
fof(f262, plain, (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1)) | ~ sP6), inference(cnf_transformation, [], [f84])).
fof(f780, plain, (~ spl18_76 | ~ spl18_8 | ~ spl18_3), inference(avatar_split_clause, [], [f263, f370, f391, f776])).
fof(f263, plain, (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2)) | ~ sP6), inference(cnf_transformation, [], [f84])).
fof(f779, plain, (~ spl18_76 | ~ spl18_4), inference(avatar_split_clause, [], [f349, f374, f776])).
fof(f349, plain, (~ (e3 = op(e3, e3)) | ~ sP6), inference(duplicate_literal_removal, [], [f264])).
fof(f264, plain, (~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | ~ sP6), inference(cnf_transformation, [], [f84])).
fof(f774, plain, (~ spl18_75 | spl18_46), inference(avatar_split_clause, [], [f253, f553, f766])).
fof(f253, plain, ((e1 = op(e1, e0)) | ~ sP7), inference(cnf_transformation, [], [f83])).
fof(f83, plain, (((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e1 = op(e0, e1)) & (e1 = op(e1, e0))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f773, plain, (~ spl18_75 | spl18_58), inference(avatar_split_clause, [], [f254, f604, f766])).
fof(f254, plain, ((e1 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f83])).
fof(f772, plain, (~ spl18_75 | ~ spl18_61), inference(avatar_split_clause, [], [f350, f617, f766])).
fof(f350, plain, (~ (e0 = op(e0, e0)) | ~ sP7), inference(duplicate_literal_removal, [], [f255])).
fof(f255, plain, (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ sP7), inference(cnf_transformation, [], [f83])).
fof(f770, plain, (~ spl18_75 | ~ spl18_53 | ~ spl18_63), inference(avatar_split_clause, [], [f257, f625, f583, f766])).
fof(f257, plain, (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2)) | ~ sP7), inference(cnf_transformation, [], [f83])).
fof(f769, plain, (~ spl18_75 | ~ spl18_49 | ~ spl18_64), inference(avatar_split_clause, [], [f258, f629, f566, f766])).
fof(f258, plain, (~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3)) | ~ sP7), inference(cnf_transformation, [], [f83])).
fof(f764, plain, (~ spl18_74 | spl18_42), inference(avatar_split_clause, [], [f247, f536, f756])).
fof(f247, plain, ((e1 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f82])).
fof(f82, plain, (((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f761, plain, (~ spl18_74 | ~ spl18_42), inference(avatar_split_clause, [], [f351, f536, f756])).
fof(f351, plain, (~ (e1 = op(e1, e1)) | ~ sP8), inference(duplicate_literal_removal, [], [f250])).
fof(f250, plain, (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f82])).
fof(f754, plain, (~ spl18_73 | spl18_38), inference(avatar_split_clause, [], [f241, f519, f746])).
fof(f241, plain, ((e1 = op(e1, e2)) | ~ sP9), inference(cnf_transformation, [], [f81])).
fof(f81, plain, (((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e1 = op(e2, e1)) & (e1 = op(e1, e2))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f753, plain, (~ spl18_73 | spl18_26), inference(avatar_split_clause, [], [f242, f468, f746])).
fof(f242, plain, ((e1 = op(e2, e1)) | ~ sP9), inference(cnf_transformation, [], [f81])).
fof(f752, plain, (~ spl18_73 | ~ spl18_31 | ~ spl18_21), inference(avatar_split_clause, [], [f243, f447, f489, f746])).
fof(f243, plain, (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP9), inference(cnf_transformation, [], [f81])).
fof(f750, plain, (~ spl18_73 | ~ spl18_23), inference(avatar_split_clause, [], [f352, f455, f746])).
fof(f352, plain, (~ (e2 = op(e2, e2)) | ~ sP9), inference(duplicate_literal_removal, [], [f245])).
fof(f245, plain, (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ sP9), inference(cnf_transformation, [], [f81])).
fof(f749, plain, (~ spl18_73 | ~ spl18_19 | ~ spl18_24), inference(avatar_split_clause, [], [f246, f459, f438, f746])).
fof(f246, plain, (~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3)) | ~ sP9), inference(cnf_transformation, [], [f81])).
fof(f744, plain, (~ spl18_72 | spl18_34), inference(avatar_split_clause, [], [f235, f502, f736])).
fof(f235, plain, ((e1 = op(e1, e3)) | ~ sP10), inference(cnf_transformation, [], [f80])).
fof(f80, plain, (((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e1 = op(e3, e1)) & (e1 = op(e1, e3))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f743, plain, (~ spl18_72 | spl18_10), inference(avatar_split_clause, [], [f236, f400, f736])).
fof(f236, plain, ((e1 = op(e3, e1)) | ~ sP10), inference(cnf_transformation, [], [f80])).
fof(f742, plain, (~ spl18_72 | ~ spl18_16 | ~ spl18_1), inference(avatar_split_clause, [], [f237, f362, f425, f736])).
fof(f237, plain, (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | ~ sP10), inference(cnf_transformation, [], [f80])).
fof(f740, plain, (~ spl18_72 | ~ spl18_8 | ~ spl18_3), inference(avatar_split_clause, [], [f239, f370, f391, f736])).
fof(f239, plain, (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2)) | ~ sP10), inference(cnf_transformation, [], [f80])).
fof(f739, plain, (~ spl18_72 | ~ spl18_4), inference(avatar_split_clause, [], [f353, f374, f736])).
fof(f353, plain, (~ (e3 = op(e3, e3)) | ~ sP10), inference(duplicate_literal_removal, [], [f240])).
fof(f240, plain, (~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | ~ sP10), inference(cnf_transformation, [], [f80])).
fof(f734, plain, (~ spl18_71 | spl18_31), inference(avatar_split_clause, [], [f229, f489, f726])).
fof(f229, plain, ((e2 = op(e2, e0)) | ~ sP11), inference(cnf_transformation, [], [f79])).
fof(f79, plain, (((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e2 = op(e0, e2)) & (e2 = op(e2, e0))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f733, plain, (~ spl18_71 | spl18_55), inference(avatar_split_clause, [], [f230, f591, f726])).
fof(f230, plain, ((e2 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f79])).
fof(f732, plain, (~ spl18_71 | ~ spl18_61), inference(avatar_split_clause, [], [f354, f617, f726])).
fof(f354, plain, (~ (e0 = op(e0, e0)) | ~ sP11), inference(duplicate_literal_removal, [], [f231])).
fof(f231, plain, (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ sP11), inference(cnf_transformation, [], [f79])).
fof(f731, plain, (~ spl18_71 | ~ spl18_57 | ~ spl18_62), inference(avatar_split_clause, [], [f232, f621, f600, f726])).
fof(f232, plain, (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1)) | ~ sP11), inference(cnf_transformation, [], [f79])).
fof(f729, plain, (~ spl18_71 | ~ spl18_49 | ~ spl18_64), inference(avatar_split_clause, [], [f234, f629, f566, f726])).
fof(f234, plain, (~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3)) | ~ sP11), inference(cnf_transformation, [], [f79])).
fof(f724, plain, (~ spl18_70 | spl18_27), inference(avatar_split_clause, [], [f223, f472, f716])).
fof(f223, plain, ((e2 = op(e2, e1)) | ~ sP12), inference(cnf_transformation, [], [f78])).
fof(f78, plain, (((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e2 = op(e1, e2)) & (e2 = op(e2, e1))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f723, plain, (~ spl18_70 | spl18_39), inference(avatar_split_clause, [], [f224, f523, f716])).
fof(f224, plain, ((e2 = op(e1, e2)) | ~ sP12), inference(cnf_transformation, [], [f78])).
fof(f722, plain, (~ spl18_70 | ~ spl18_46 | ~ spl18_41), inference(avatar_split_clause, [], [f225, f532, f553, f716])).
fof(f225, plain, (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0)) | ~ sP12), inference(cnf_transformation, [], [f78])).
fof(f721, plain, (~ spl18_70 | ~ spl18_42), inference(avatar_split_clause, [], [f355, f536, f716])).
fof(f355, plain, (~ (e1 = op(e1, e1)) | ~ sP12), inference(duplicate_literal_removal, [], [f226])).
fof(f226, plain, (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ sP12), inference(cnf_transformation, [], [f78])).
fof(f720, plain, (~ spl18_70 | ~ spl18_38 | ~ spl18_43), inference(avatar_split_clause, [], [f227, f540, f519, f716])).
fof(f227, plain, (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2)) | ~ sP12), inference(cnf_transformation, [], [f78])).
fof(f719, plain, (~ spl18_70 | ~ spl18_34 | ~ spl18_44), inference(avatar_split_clause, [], [f228, f544, f502, f716])).
fof(f228, plain, (~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3)) | ~ sP12), inference(cnf_transformation, [], [f78])).
fof(f714, plain, (~ spl18_69 | spl18_23), inference(avatar_split_clause, [], [f217, f455, f706])).
fof(f217, plain, ((e2 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f77])).
fof(f77, plain, (((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f710, plain, (~ spl18_69 | ~ spl18_23), inference(avatar_split_clause, [], [f356, f455, f706])).
fof(f356, plain, (~ (e2 = op(e2, e2)) | ~ sP13), inference(duplicate_literal_removal, [], [f221])).
fof(f221, plain, (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f77])).
fof(f704, plain, (~ spl18_68 | spl18_19), inference(avatar_split_clause, [], [f211, f438, f696])).
fof(f211, plain, ((e2 = op(e2, e3)) | ~ sP14), inference(cnf_transformation, [], [f76])).
fof(f76, plain, (((~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & (~ (e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & (e2 = op(e3, e2)) & (e2 = op(e2, e3))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f703, plain, (~ spl18_68 | spl18_7), inference(avatar_split_clause, [], [f212, f387, f696])).
fof(f212, plain, ((e2 = op(e3, e2)) | ~ sP14), inference(cnf_transformation, [], [f76])).
fof(f702, plain, (~ spl18_68 | ~ spl18_16 | ~ spl18_1), inference(avatar_split_clause, [], [f213, f362, f425, f696])).
fof(f213, plain, (~ (e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | ~ sP14), inference(cnf_transformation, [], [f76])).
fof(f701, plain, (~ spl18_68 | ~ spl18_12 | ~ spl18_2), inference(avatar_split_clause, [], [f214, f366, f408, f696])).
fof(f214, plain, (~ (e1 = op(e3, e3)) | ~ (e3 = op(e3, e1)) | ~ sP14), inference(cnf_transformation, [], [f76])).
fof(f699, plain, (~ spl18_68 | ~ spl18_4), inference(avatar_split_clause, [], [f357, f374, f696])).
fof(f357, plain, (~ (e3 = op(e3, e3)) | ~ sP14), inference(duplicate_literal_removal, [], [f216])).
fof(f216, plain, (~ (e3 = op(e3, e3)) | ~ (e3 = op(e3, e3)) | ~ sP14), inference(cnf_transformation, [], [f76])).
fof(f694, plain, (~ spl18_67 | spl18_16), inference(avatar_split_clause, [], [f205, f425, f686])).
fof(f205, plain, ((e3 = op(e3, e0)) | ~ sP15), inference(cnf_transformation, [], [f75])).
fof(f75, plain, (((~ (op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e3 = op(e0, e3)) & (e3 = op(e3, e0))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f693, plain, (~ spl18_67 | spl18_52), inference(avatar_split_clause, [], [f206, f578, f686])).
fof(f206, plain, ((e3 = op(e0, e3)) | ~ sP15), inference(cnf_transformation, [], [f75])).
fof(f692, plain, (~ spl18_67 | ~ spl18_61), inference(avatar_split_clause, [], [f358, f617, f686])).
fof(f358, plain, (~ (e0 = op(e0, e0)) | ~ sP15), inference(duplicate_literal_removal, [], [f207])).
fof(f207, plain, (~ (e0 = op(e0, e0)) | ~ (e0 = op(e0, e0)) | ~ sP15), inference(cnf_transformation, [], [f75])).
fof(f691, plain, (~ spl18_67 | ~ spl18_57 | ~ spl18_62), inference(avatar_split_clause, [], [f208, f621, f600, f686])).
fof(f208, plain, (~ (op(e0, e0) = e1) | ~ (e0 = op(e0, e1)) | ~ sP15), inference(cnf_transformation, [], [f75])).
fof(f690, plain, (~ spl18_67 | ~ spl18_53 | ~ spl18_63), inference(avatar_split_clause, [], [f209, f625, f583, f686])).
fof(f209, plain, (~ (op(e0, e0) = e2) | ~ (e0 = op(e0, e2)) | ~ sP15), inference(cnf_transformation, [], [f75])).
fof(f684, plain, (~ spl18_66 | spl18_12), inference(avatar_split_clause, [], [f199, f408, f676])).
fof(f199, plain, ((e3 = op(e3, e1)) | ~ sP16), inference(cnf_transformation, [], [f74])).
fof(f74, plain, (((~ (e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & (e3 = op(e1, e3)) & (e3 = op(e3, e1))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f683, plain, (~ spl18_66 | spl18_36), inference(avatar_split_clause, [], [f200, f510, f676])).
fof(f200, plain, ((e3 = op(e1, e3)) | ~ sP16), inference(cnf_transformation, [], [f74])).
fof(f682, plain, (~ spl18_66 | ~ spl18_46 | ~ spl18_41), inference(avatar_split_clause, [], [f201, f532, f553, f676])).
fof(f201, plain, (~ (e0 = op(e1, e1)) | ~ (e1 = op(e1, e0)) | ~ sP16), inference(cnf_transformation, [], [f74])).
fof(f681, plain, (~ spl18_66 | ~ spl18_42), inference(avatar_split_clause, [], [f359, f536, f676])).
fof(f359, plain, (~ (e1 = op(e1, e1)) | ~ sP16), inference(duplicate_literal_removal, [], [f202])).
fof(f202, plain, (~ (e1 = op(e1, e1)) | ~ (e1 = op(e1, e1)) | ~ sP16), inference(cnf_transformation, [], [f74])).
fof(f680, plain, (~ spl18_66 | ~ spl18_38 | ~ spl18_43), inference(avatar_split_clause, [], [f203, f540, f519, f676])).
fof(f203, plain, (~ (e2 = op(e1, e1)) | ~ (e1 = op(e1, e2)) | ~ sP16), inference(cnf_transformation, [], [f74])).
fof(f674, plain, (~ spl18_65 | spl18_8), inference(avatar_split_clause, [], [f193, f391, f666])).
fof(f193, plain, ((e3 = op(e3, e2)) | ~ sP17), inference(cnf_transformation, [], [f73])).
fof(f73, plain, (((~ (e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & (e3 = op(e2, e3)) & (e3 = op(e3, e2))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f673, plain, (~ spl18_65 | spl18_20), inference(avatar_split_clause, [], [f194, f442, f666])).
fof(f194, plain, ((e3 = op(e2, e3)) | ~ sP17), inference(cnf_transformation, [], [f73])).
fof(f672, plain, (~ spl18_65 | ~ spl18_31 | ~ spl18_21), inference(avatar_split_clause, [], [f195, f447, f489, f666])).
fof(f195, plain, (~ (e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP17), inference(cnf_transformation, [], [f73])).
fof(f671, plain, (~ spl18_65 | ~ spl18_27 | ~ spl18_22), inference(avatar_split_clause, [], [f196, f451, f472, f666])).
fof(f196, plain, (~ (e1 = op(e2, e2)) | ~ (e2 = op(e2, e1)) | ~ sP17), inference(cnf_transformation, [], [f73])).
fof(f670, plain, (~ spl18_65 | ~ spl18_23), inference(avatar_split_clause, [], [f360, f455, f666])).
fof(f360, plain, (~ (e2 = op(e2, e2)) | ~ sP17), inference(duplicate_literal_removal, [], [f197])).
fof(f197, plain, (~ (e2 = op(e2, e2)) | ~ (e2 = op(e2, e2)) | ~ sP17), inference(cnf_transformation, [], [f73])).
fof(f664, plain, (spl18_61 | spl18_57 | spl18_53 | spl18_49), inference(avatar_split_clause, [], [f107, f566, f583, f600, f617])).
fof(f107, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax2)).
fof(f663, plain, (spl18_61 | spl18_45 | spl18_29 | spl18_13), inference(avatar_split_clause, [], [f108, f413, f481, f549, f617])).
fof(f108, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f662, plain, (spl18_62 | spl18_58 | spl18_54 | spl18_50), inference(avatar_split_clause, [], [f109, f570, f587, f604, f621])).
fof(f109, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f661, plain, (spl18_62 | spl18_46 | spl18_30 | spl18_14), inference(avatar_split_clause, [], [f110, f417, f485, f553, f621])).
fof(f110, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f660, plain, (spl18_63 | spl18_59 | spl18_55 | spl18_51), inference(avatar_split_clause, [], [f111, f574, f591, f608, f625])).
fof(f111, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f659, plain, (spl18_63 | spl18_47 | spl18_31 | spl18_15), inference(avatar_split_clause, [], [f112, f421, f489, f557, f625])).
fof(f112, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f658, plain, (spl18_64 | spl18_60 | spl18_56 | spl18_52), inference(avatar_split_clause, [], [f113, f578, f595, f612, f629])).
fof(f113, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f657, plain, (spl18_64 | spl18_48 | spl18_32 | spl18_16), inference(avatar_split_clause, [], [f114, f425, f493, f561, f629])).
fof(f114, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f656, plain, (spl18_45 | spl18_41 | spl18_37 | spl18_33), inference(avatar_split_clause, [], [f115, f498, f515, f532, f549])).
fof(f115, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f655, plain, (spl18_57 | spl18_41 | spl18_25 | spl18_9), inference(avatar_split_clause, [], [f116, f396, f464, f532, f600])).
fof(f116, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f654, plain, (spl18_46 | spl18_42 | spl18_38 | spl18_34), inference(avatar_split_clause, [], [f117, f502, f519, f536, f553])).
fof(f117, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f653, plain, (spl18_58 | spl18_42 | spl18_26 | spl18_10), inference(avatar_split_clause, [], [f118, f400, f468, f536, f604])).
fof(f118, plain, ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f652, plain, (spl18_47 | spl18_43 | spl18_39 | spl18_35), inference(avatar_split_clause, [], [f119, f506, f523, f540, f557])).
fof(f119, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f651, plain, (spl18_59 | spl18_43 | spl18_27 | spl18_11), inference(avatar_split_clause, [], [f120, f404, f472, f540, f608])).
fof(f120, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f650, plain, (spl18_48 | spl18_44 | spl18_40 | spl18_36), inference(avatar_split_clause, [], [f121, f510, f527, f544, f561])).
fof(f121, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f649, plain, (spl18_60 | spl18_44 | spl18_28 | spl18_12), inference(avatar_split_clause, [], [f122, f408, f476, f544, f612])).
fof(f122, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f648, plain, (spl18_29 | spl18_25 | spl18_21 | spl18_17), inference(avatar_split_clause, [], [f123, f430, f447, f464, f481])).
fof(f123, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f647, plain, (spl18_53 | spl18_37 | spl18_21 | spl18_5), inference(avatar_split_clause, [], [f124, f379, f447, f515, f583])).
fof(f124, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f646, plain, (spl18_30 | spl18_26 | spl18_22 | spl18_18), inference(avatar_split_clause, [], [f125, f434, f451, f468, f485])).
fof(f125, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f645, plain, (spl18_54 | spl18_38 | spl18_22 | spl18_6), inference(avatar_split_clause, [], [f126, f383, f451, f519, f587])).
fof(f126, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f644, plain, (spl18_31 | spl18_27 | spl18_23 | spl18_19), inference(avatar_split_clause, [], [f127, f438, f455, f472, f489])).
fof(f127, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f643, plain, (spl18_55 | spl18_39 | spl18_23 | spl18_7), inference(avatar_split_clause, [], [f128, f387, f455, f523, f591])).
fof(f128, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f642, plain, (spl18_32 | spl18_28 | spl18_24 | spl18_20), inference(avatar_split_clause, [], [f129, f442, f459, f476, f493])).
fof(f129, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f641, plain, (spl18_56 | spl18_40 | spl18_24 | spl18_8), inference(avatar_split_clause, [], [f130, f391, f459, f527, f595])).
fof(f130, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f640, plain, (spl18_13 | spl18_9 | spl18_5 | spl18_1), inference(avatar_split_clause, [], [f131, f362, f379, f396, f413])).
fof(f131, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f639, plain, (spl18_49 | spl18_33 | spl18_17 | spl18_1), inference(avatar_split_clause, [], [f132, f362, f430, f498, f566])).
fof(f132, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f638, plain, (spl18_14 | spl18_10 | spl18_6 | spl18_2), inference(avatar_split_clause, [], [f133, f366, f383, f400, f417])).
fof(f133, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f637, plain, (spl18_50 | spl18_34 | spl18_18 | spl18_2), inference(avatar_split_clause, [], [f134, f366, f434, f502, f570])).
fof(f134, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f636, plain, (spl18_15 | spl18_11 | spl18_7 | spl18_3), inference(avatar_split_clause, [], [f135, f370, f387, f404, f421])).
fof(f135, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f635, plain, (spl18_51 | spl18_35 | spl18_19 | spl18_3), inference(avatar_split_clause, [], [f136, f370, f438, f506, f574])).
fof(f136, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f634, plain, (spl18_16 | spl18_12 | spl18_8 | spl18_4), inference(avatar_split_clause, [], [f137, f374, f391, f408, f425])).
fof(f137, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f633, plain, (spl18_52 | spl18_36 | spl18_20 | spl18_4), inference(avatar_split_clause, [], [f138, f374, f442, f510, f578])).
fof(f138, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f632, plain, (spl18_61 | spl18_62 | spl18_63 | spl18_64), inference(avatar_split_clause, [], [f91, f629, f625, f621, f617])).
fof(f91, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG151+1.p', ax1)).
fof(f615, plain, (spl18_57 | spl18_58 | spl18_59 | spl18_60), inference(avatar_split_clause, [], [f92, f612, f608, f604, f600])).
fof(f92, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f598, plain, (spl18_53 | spl18_54 | spl18_55 | spl18_56), inference(avatar_split_clause, [], [f93, f595, f591, f587, f583])).
fof(f93, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f581, plain, (spl18_49 | spl18_50 | spl18_51 | spl18_52), inference(avatar_split_clause, [], [f94, f578, f574, f570, f566])).
fof(f94, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f564, plain, (spl18_45 | spl18_46 | spl18_47 | spl18_48), inference(avatar_split_clause, [], [f95, f561, f557, f553, f549])).
fof(f95, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f547, plain, (spl18_41 | spl18_42 | spl18_43 | spl18_44), inference(avatar_split_clause, [], [f96, f544, f540, f536, f532])).
fof(f96, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f530, plain, (spl18_37 | spl18_38 | spl18_39 | spl18_40), inference(avatar_split_clause, [], [f97, f527, f523, f519, f515])).
fof(f97, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f513, plain, (spl18_33 | spl18_34 | spl18_35 | spl18_36), inference(avatar_split_clause, [], [f98, f510, f506, f502, f498])).
fof(f98, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f496, plain, (spl18_29 | spl18_30 | spl18_31 | spl18_32), inference(avatar_split_clause, [], [f99, f493, f489, f485, f481])).
fof(f99, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f479, plain, (spl18_25 | spl18_26 | spl18_27 | spl18_28), inference(avatar_split_clause, [], [f100, f476, f472, f468, f464])).
fof(f100, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f462, plain, (spl18_21 | spl18_22 | spl18_23 | spl18_24), inference(avatar_split_clause, [], [f101, f459, f455, f451, f447])).
fof(f101, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f445, plain, (spl18_17 | spl18_18 | spl18_19 | spl18_20), inference(avatar_split_clause, [], [f102, f442, f438, f434, f430])).
fof(f102, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f428, plain, (spl18_13 | spl18_14 | spl18_15 | spl18_16), inference(avatar_split_clause, [], [f103, f425, f421, f417, f413])).
fof(f103, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f411, plain, (spl18_9 | spl18_10 | spl18_11 | spl18_12), inference(avatar_split_clause, [], [f104, f408, f404, f400, f396])).
fof(f104, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f394, plain, (spl18_5 | spl18_6 | spl18_7 | spl18_8), inference(avatar_split_clause, [], [f105, f391, f387, f383, f379])).
fof(f105, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f377, plain, (spl18_1 | spl18_2 | spl18_3 | spl18_4), inference(avatar_split_clause, [], [f106, f374, f370, f366, f362])).
fof(f106, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).