fof(f3674, plain, $false, inference(avatar_sat_refutation, [], [f605, f622, f639, f656, f673, f690, f707, f724, f741, f758, f775, f792, f809, f826, f843, f860, f861, f862, f863, f864, f865, f866, f867, f869, f870, f872, f873, f874, f875, f876, f877, f878, f880, f883, f884, f885, f886, f887, f888, f889, f890, f892, f897, f898, f899, f900, f905, f906, f907, f908, f913, f914, f915, f916, f923, f924, f931, f932, f939, f940, f946, f948, f953, f954, f956, f961, f962, f963, f964, f969, f970, f971, f972, f978, f980, f985, f986, f987, f988, f994, f996, f1001, f1002, f1003, f1004, f1010, f1012, f1017, f1018, f1019, f1020, f1025, f1026, f1027, f1028, f1034, f1036, f1042, f1044, f1050, f1052, f1057, f1058, f1059, f1060, f1065, f1066, f1067, f1068, f1073, f1075, f1076, f1083, f1084, f1091, f1092, f1098, f1099, f1100, f1105, f1106, f1107, f1108, f1114, f1115, f1116, f1122, f1124, f1129, f1130, f1131, f1132, f1137, f1138, f1139, f1140, f1146, f1148, f1153, f1154, f1155, f1156, f1162, f1164, f1170, f1172, f1177, f1178, f1179, f1180, f1186, f1188, f1193, f1194, f1195, f1196, f1201, f1202, f1203, f1204, f1209, f1210, f1212, f1218, f1220, f1225, f1226, f1227, f1228, f1235, f1236, f1243, f1244, f1251, f1252, f1258, f1259, f1260, f1265, f1266, f1267, f1268, f1273, f1274, f1275, f1276, f1282, f1284, f1290, f1291, f1292, f1298, f1299, f1300, f1305, f1306, f1307, f1308, f1313, f1314, f1315, f1316, f1322, f1324, f1329, f1330, f1331, f1332, f1337, f1338, f1340, f1345, f1346, f1347, f1348, f1354, f1356, f1361, f1362, f1363, f1364, f1369, f1370, f1371, f1372, f1378, f1380, f1386, f1388, f1393, f1395, f1396, f1401, f1403, f1404, f1409, f1411, f1412, f1419, f1420, f1421, f1422, f1423, f1424, f1427, f1428, f1437, f1442, f1451, f1456, f1465, f1470, f1479, f1484, f1489, f1494, f1499, f1509, f1514, f1519, f1524, f1529, f1530, f1535, f1541, f1559, f1562, f1565, f1566, f1569, f1575, f1578, f1596, f1616, f1618, f1628, f1635, f1652, f1659, f1660, f1663, f1665, f1668, f1677, f1687, f1693, f1720, f1760, f1778, f1782, f1798, f1804, f1810, f1815, f1837, f1843, f1875, f1882, f1900, f1915, f1918, f1929, f1955, f1968, f1979, f1984, f2003, f2020, f2034, f2042, f2044, f2137, f2162, f2175, f2184, f2187, f2200, f2208, f2233, f2244, f2254, f2258, f2266, f2271, f2277, f2283, f2289, f2296, f2304, f2307, f2353, f2365, f2378, f2381, f2391, f2395, f2404, f2407, f2413, f2419, f2443, f2454, f2458, f2463, f2480, f2531, f2561, f2565, f2587, f2595, f2610, f2616, f2619, f2627, f2628, f2633, f2665, f2668, f2677, f2680, f2691, f2708, f2710, f2711, f2717, f2723, f2735, f2768, f2780, f2804, f2813, f2820, f2827, f2830, f2851, f2860, f2863, f2870, f2878, f2879, f2884, f2885, f2892, f2893, f2904, f2915, f2921, f2929, f2931, f2937, f2967, f2968, f2981, f2987, f2990, f3026, f3027, f3028, f3037, f3042, f3044, f3054, f3058, f3061, f3085, f3086, f3091, f3097, f3098, f3107, f3108, f3118, f3122, f3131, f3148, f3150, f3157, f3172, f3186, f3193, f3203, f3204, f3211, f3217, f3222, f3224, f3244, f3245, f3255, f3261, f3262, f3272, f3285, f3289, f3308, f3327, f3346, f3347, f3355, f3356, f3368, f3376, f3385, f3398, f3403, f3405, f3410, f3419, f3424, f3426, f3444, f3445, f3448, f3451, f3456, f3463, f3470, f3475, f3478, f3495, f3501, f3503, f3506, f3508, f3510, f3513, f3515, f3517, f3529, f3536, f3546, f3553, f3554, f3559, f3560, f3565, f3570, f3571, f3576, f3587, f3588, f3592, f3624, f3638, f3639, f3645, f3651, f3661, f3664, f3673])).
fof(f3673, plain, (~ spl66_27 | ~ spl66_31), inference(avatar_split_clause, [], [f3671, f717, f700])).
fof(f700, plain, (spl66_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_27])])).
fof(f717, plain, (spl66_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_31])])).
fof(f3671, plain, (~ (e2 = op(e2, e1)) | ~ spl66_31), inference(backward_demodulation, [], [f271, f719])).
fof(f719, plain, ((e2 = op(e2, e0)) | ~ spl66_31), inference(avatar_component_clause, [], [f717])).
fof(f271, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax3)).
fof(f3664, plain, (~ spl66_21 | ~ spl66_43 | spl66_131), inference(avatar_contradiction_clause, [], [f3663])).
fof(f3663, plain, ($false | (~ spl66_21 | ~ spl66_43 | spl66_131)), inference(subsumption_resolution, [], [f3655, f677])).
fof(f677, plain, ((e0 = op(e2, e2)) | ~ spl66_21), inference(avatar_component_clause, [], [f675])).
fof(f675, plain, (spl66_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_21])])).
fof(f3655, plain, (~ (e0 = op(e2, e2)) | (~ spl66_43 | spl66_131)), inference(backward_demodulation, [], [f1432, f770])).
fof(f770, plain, ((e2 = op(e1, e1)) | ~ spl66_43), inference(avatar_component_clause, [], [f768])).
fof(f768, plain, (spl66_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_43])])).
fof(f1432, plain, (~ (e0 = op(op(e1, e1), op(e1, e1))) | spl66_131), inference(avatar_component_clause, [], [f1430])).
fof(f1430, plain, (spl66_131 <=> (e0 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_131])])).
fof(f3661, plain, (~ spl66_27 | ~ spl66_43), inference(avatar_split_clause, [], [f3652, f768, f700])).
fof(f3652, plain, (~ (e2 = op(e2, e1)) | ~ spl66_43), inference(backward_demodulation, [], [f244, f770])).
fof(f244, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3651, plain, (~ spl66_38 | ~ spl66_46), inference(avatar_split_clause, [], [f3648, f781, f747])).
fof(f747, plain, (spl66_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_38])])).
fof(f781, plain, (spl66_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_46])])).
fof(f3648, plain, (~ (e1 = op(e1, e2)) | ~ spl66_46), inference(backward_demodulation, [], [f266, f783])).
fof(f783, plain, ((e1 = op(e1, e0)) | ~ spl66_46), inference(avatar_component_clause, [], [f781])).
fof(f266, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3645, plain, (~ spl66_38 | ~ spl66_54), inference(avatar_split_clause, [], [f3642, f815, f747])).
fof(f815, plain, (spl66_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_54])])).
fof(f3642, plain, (~ (e1 = op(e1, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f247, f817])).
fof(f817, plain, ((e1 = op(e0, e2)) | ~ spl66_54), inference(avatar_component_clause, [], [f815])).
fof(f247, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3639, plain, (~ spl66_56 | ~ spl66_64), inference(avatar_split_clause, [], [f3633, f857, f823])).
fof(f823, plain, (spl66_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_56])])).
fof(f857, plain, (spl66_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl66_64])])).
fof(f3633, plain, (~ (e3 = op(e0, e2)) | ~ spl66_64), inference(backward_demodulation, [], [f260, f859])).
fof(f859, plain, ((op(e0, e0) = e3) | ~ spl66_64), inference(avatar_component_clause, [], [f857])).
fof(f260, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3638, plain, (~ spl66_32 | ~ spl66_64), inference(avatar_split_clause, [], [f3632, f857, f721])).
fof(f721, plain, (spl66_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_32])])).
fof(f3632, plain, (~ (e3 = op(e2, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f236, f859])).
fof(f236, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3624, plain, (~ spl66_56 | ~ spl66_21 | spl66_135), inference(avatar_split_clause, [], [f3582, f1448, f675, f823])).
fof(f1448, plain, (spl66_135 <=> (e3 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_135])])).
fof(f3582, plain, (~ (e3 = op(e0, e2)) | (~ spl66_21 | spl66_135)), inference(backward_demodulation, [], [f1450, f677])).
fof(f1450, plain, (~ (e3 = op(op(e2, e2), e2)) | spl66_135), inference(avatar_component_clause, [], [f1448])).
fof(f3592, plain, (~ spl66_21 | ~ spl66_62 | spl66_143), inference(avatar_contradiction_clause, [], [f3591])).
fof(f3591, plain, ($false | (~ spl66_21 | ~ spl66_62 | spl66_143)), inference(subsumption_resolution, [], [f3584, f851])).
fof(f851, plain, ((op(e0, e0) = e1) | ~ spl66_62), inference(avatar_component_clause, [], [f849])).
fof(f849, plain, (spl66_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl66_62])])).
fof(f3584, plain, (~ (op(e0, e0) = e1) | (~ spl66_21 | spl66_143)), inference(backward_demodulation, [], [f1488, f677])).
fof(f1488, plain, (~ (e1 = op(op(e2, e2), op(e2, e2))) | spl66_143), inference(avatar_component_clause, [], [f1486])).
fof(f1486, plain, (spl66_143 <=> (e1 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_143])])).
fof(f3588, plain, (~ spl66_17 | ~ spl66_21), inference(avatar_split_clause, [], [f3578, f675, f658])).
fof(f658, plain, (spl66_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_17])])).
fof(f3578, plain, (~ (e0 = op(e2, e3)) | ~ spl66_21), inference(backward_demodulation, [], [f276, f677])).
fof(f276, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3587, plain, (~ spl66_5 | ~ spl66_21), inference(avatar_split_clause, [], [f3577, f675, f607])).
fof(f607, plain, (spl66_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_5])])).
fof(f3577, plain, (~ (e0 = op(e3, e2)) | ~ spl66_21), inference(backward_demodulation, [], [f252, f677])).
fof(f252, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3576, plain, (~ spl66_11 | ~ spl66_27), inference(avatar_split_clause, [], [f3572, f700, f632])).
fof(f632, plain, (spl66_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_11])])).
fof(f3572, plain, (~ (e2 = op(e3, e1)) | ~ spl66_27), inference(backward_demodulation, [], [f246, f702])).
fof(f702, plain, ((e2 = op(e2, e1)) | ~ spl66_27), inference(avatar_component_clause, [], [f700])).
fof(f246, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3571, plain, (~ spl66_28 | ~ spl66_32), inference(avatar_split_clause, [], [f3567, f721, f704])).
fof(f704, plain, (spl66_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_28])])).
fof(f3567, plain, (~ (e3 = op(e2, e1)) | ~ spl66_32), inference(backward_demodulation, [], [f271, f723])).
fof(f723, plain, ((e3 = op(e2, e0)) | ~ spl66_32), inference(avatar_component_clause, [], [f721])).
fof(f3570, plain, (~ spl66_16 | ~ spl66_32), inference(avatar_split_clause, [], [f3566, f721, f653])).
fof(f653, plain, (spl66_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_16])])).
fof(f3566, plain, (~ (e3 = op(e3, e0)) | ~ spl66_32), inference(backward_demodulation, [], [f240, f723])).
fof(f240, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3565, plain, (~ spl66_17 | ~ spl66_33), inference(avatar_split_clause, [], [f3563, f726, f658])).
fof(f726, plain, (spl66_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_33])])).
fof(f3563, plain, (~ (e0 = op(e2, e3)) | ~ spl66_33), inference(backward_demodulation, [], [f256, f728])).
fof(f728, plain, ((e0 = op(e1, e3)) | ~ spl66_33), inference(avatar_component_clause, [], [f726])).
fof(f256, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3560, plain, (~ spl66_34 | ~ spl66_38), inference(avatar_split_clause, [], [f3557, f747, f730])).
fof(f730, plain, (spl66_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_34])])).
fof(f3557, plain, (~ (e1 = op(e1, e3)) | ~ spl66_38), inference(backward_demodulation, [], [f270, f749])).
fof(f749, plain, ((e1 = op(e1, e2)) | ~ spl66_38), inference(avatar_component_clause, [], [f747])).
fof(f270, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3559, plain, (~ spl66_22 | ~ spl66_38), inference(avatar_split_clause, [], [f3555, f747, f679])).
fof(f679, plain, (spl66_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_22])])).
fof(f3555, plain, (~ (e1 = op(e2, e2)) | ~ spl66_38), inference(backward_demodulation, [], [f250, f749])).
fof(f250, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3554, plain, (~ spl66_39 | ~ spl66_47), inference(avatar_split_clause, [], [f3550, f785, f751])).
fof(f751, plain, (spl66_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_39])])).
fof(f785, plain, (spl66_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_47])])).
fof(f3550, plain, (~ (e2 = op(e1, e2)) | ~ spl66_47), inference(backward_demodulation, [], [f266, f787])).
fof(f787, plain, ((e2 = op(e1, e0)) | ~ spl66_47), inference(avatar_component_clause, [], [f785])).
fof(f3553, plain, (~ spl66_31 | ~ spl66_47), inference(avatar_split_clause, [], [f3548, f785, f717])).
fof(f3548, plain, (~ (e2 = op(e2, e0)) | ~ spl66_47), inference(backward_demodulation, [], [f238, f787])).
fof(f238, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3546, plain, (~ spl66_47 | ~ spl66_62 | spl66_142), inference(avatar_split_clause, [], [f3543, f1481, f849, f785])).
fof(f1481, plain, (spl66_142 <=> (e2 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_142])])).
fof(f3543, plain, (~ (e2 = op(e1, e0)) | (~ spl66_62 | spl66_142)), inference(backward_demodulation, [], [f1483, f851])).
fof(f1483, plain, (~ (e2 = op(op(e0, e0), e0)) | spl66_142), inference(avatar_component_clause, [], [f1481])).
fof(f3536, plain, (~ spl66_11 | ~ spl66_44 | spl66_133), inference(avatar_split_clause, [], [f3535, f1439, f772, f632])).
fof(f772, plain, (spl66_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_44])])).
fof(f1439, plain, (spl66_133 <=> (e2 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_133])])).
fof(f3535, plain, (~ (e2 = op(e3, e1)) | (~ spl66_44 | spl66_133)), inference(forward_demodulation, [], [f1441, f774])).
fof(f774, plain, ((e3 = op(e1, e1)) | ~ spl66_44), inference(avatar_component_clause, [], [f772])).
fof(f1441, plain, (~ (e2 = op(op(e1, e1), e1)) | spl66_133), inference(avatar_component_clause, [], [f1439])).
fof(f3529, plain, (~ spl66_63 | ~ spl66_51), inference(avatar_split_clause, [], [f3528, f802, f853])).
fof(f853, plain, (spl66_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl66_63])])).
fof(f802, plain, (spl66_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_51])])).
fof(f3528, plain, (~ (op(e0, e0) = e2) | ~ spl66_51), inference(forward_demodulation, [], [f261, f804])).
fof(f804, plain, ((e2 = op(e0, e3)) | ~ spl66_51), inference(avatar_component_clause, [], [f802])).
fof(f261, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3517, plain, (~ spl66_40 | ~ spl66_44), inference(avatar_split_clause, [], [f3516, f772, f755])).
fof(f755, plain, (spl66_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_40])])).
fof(f3516, plain, (~ (e3 = op(e1, e2)) | ~ spl66_44), inference(forward_demodulation, [], [f268, f774])).
fof(f268, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3515, plain, (~ spl66_40 | ~ spl66_56), inference(avatar_split_clause, [], [f3514, f823, f755])).
fof(f3514, plain, (~ (e3 = op(e1, e2)) | ~ spl66_56), inference(forward_demodulation, [], [f247, f825])).
fof(f825, plain, ((e3 = op(e0, e2)) | ~ spl66_56), inference(avatar_component_clause, [], [f823])).
fof(f3513, plain, (~ spl66_36 | ~ spl66_4), inference(avatar_split_clause, [], [f3512, f602, f738])).
fof(f738, plain, (spl66_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_36])])).
fof(f602, plain, (spl66_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_4])])).
fof(f3512, plain, (~ (e3 = op(e1, e3)) | ~ spl66_4), inference(forward_demodulation, [], [f257, f604])).
fof(f604, plain, ((e3 = op(e3, e3)) | ~ spl66_4), inference(avatar_component_clause, [], [f602])).
fof(f257, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3510, plain, (~ spl66_16 | ~ spl66_4), inference(avatar_split_clause, [], [f3509, f602, f653])).
fof(f3509, plain, (~ (e3 = op(e3, e0)) | ~ spl66_4), inference(forward_demodulation, [], [f279, f604])).
fof(f279, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3508, plain, (~ spl66_12 | ~ spl66_4), inference(avatar_split_clause, [], [f3507, f602, f636])).
fof(f636, plain, (spl66_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_12])])).
fof(f3507, plain, (~ (e3 = op(e3, e1)) | ~ spl66_4), inference(forward_demodulation, [], [f281, f604])).
fof(f281, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3506, plain, (~ spl66_12 | ~ spl66_44), inference(avatar_split_clause, [], [f3505, f772, f636])).
fof(f3505, plain, (~ (e3 = op(e3, e1)) | ~ spl66_44), inference(forward_demodulation, [], [f245, f774])).
fof(f245, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3503, plain, (~ spl66_8 | ~ spl66_4), inference(avatar_split_clause, [], [f3502, f602, f619])).
fof(f619, plain, (spl66_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_8])])).
fof(f3502, plain, (~ (e3 = op(e3, e2)) | ~ spl66_4), inference(forward_demodulation, [], [f282, f604])).
fof(f282, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3501, plain, (~ spl66_8 | ~ spl66_56), inference(avatar_split_clause, [], [f3500, f823, f619])).
fof(f3500, plain, (~ (e3 = op(e3, e2)) | ~ spl66_56), inference(forward_demodulation, [], [f249, f825])).
fof(f249, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3495, plain, (~ spl66_14 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f3494])).
fof(f3494, plain, ($false | (~ spl66_14 | ~ spl66_16)), inference(subsumption_resolution, [], [f3493, f287])).
fof(f287, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax4)).
fof(f3493, plain, ((e1 = e3) | (~ spl66_14 | ~ spl66_16)), inference(backward_demodulation, [], [f655, f647])).
fof(f647, plain, ((e1 = op(e3, e0)) | ~ spl66_14), inference(avatar_component_clause, [], [f645])).
fof(f645, plain, (spl66_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_14])])).
fof(f655, plain, ((e3 = op(e3, e0)) | ~ spl66_16), inference(avatar_component_clause, [], [f653])).
fof(f3478, plain, (~ spl66_24 | ~ spl66_28), inference(avatar_split_clause, [], [f3476, f704, f687])).
fof(f687, plain, (spl66_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_24])])).
fof(f3476, plain, (~ (e3 = op(e2, e2)) | ~ spl66_28), inference(backward_demodulation, [], [f274, f706])).
fof(f706, plain, ((e3 = op(e2, e1)) | ~ spl66_28), inference(avatar_component_clause, [], [f704])).
fof(f274, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3475, plain, (~ spl66_19 | ~ spl66_31), inference(avatar_split_clause, [], [f3472, f717, f666])).
fof(f666, plain, (spl66_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_19])])).
fof(f3472, plain, (~ (e2 = op(e2, e3)) | ~ spl66_31), inference(backward_demodulation, [], [f273, f719])).
fof(f273, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3470, plain, (~ spl66_2 | ~ spl66_34), inference(avatar_split_clause, [], [f3468, f730, f594])).
fof(f594, plain, (spl66_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_2])])).
fof(f3468, plain, (~ (e1 = op(e3, e3)) | ~ spl66_34), inference(backward_demodulation, [], [f257, f732])).
fof(f732, plain, ((e1 = op(e1, e3)) | ~ spl66_34), inference(avatar_component_clause, [], [f730])).
fof(f3463, plain, (~ spl66_28 | ~ spl66_44), inference(avatar_split_clause, [], [f3457, f772, f704])).
fof(f3457, plain, (~ (e3 = op(e2, e1)) | ~ spl66_44), inference(backward_demodulation, [], [f244, f774])).
fof(f3456, plain, (~ spl66_19 | ~ spl66_51), inference(avatar_split_clause, [], [f3453, f802, f666])).
fof(f3453, plain, (~ (e2 = op(e2, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f254, f804])).
fof(f254, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3451, plain, (~ spl66_24 | ~ spl66_56), inference(avatar_split_clause, [], [f3449, f823, f687])).
fof(f3449, plain, (~ (e3 = op(e2, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f248, f825])).
fof(f248, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3448, plain, (~ spl66_57 | ~ spl66_58), inference(avatar_contradiction_clause, [], [f3447])).
fof(f3447, plain, ($false | (~ spl66_57 | ~ spl66_58)), inference(subsumption_resolution, [], [f3446, f283])).
fof(f283, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f3446, plain, ((e0 = e1) | (~ spl66_57 | ~ spl66_58)), inference(backward_demodulation, [], [f834, f830])).
fof(f830, plain, ((e0 = op(e0, e1)) | ~ spl66_57), inference(avatar_component_clause, [], [f828])).
fof(f828, plain, (spl66_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_57])])).
fof(f834, plain, ((e1 = op(e0, e1)) | ~ spl66_58), inference(avatar_component_clause, [], [f832])).
fof(f832, plain, (spl66_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_58])])).
fof(f3445, plain, (~ spl66_54 | ~ spl66_58), inference(avatar_split_clause, [], [f3440, f832, f815])).
fof(f3440, plain, (~ (e1 = op(e0, e2)) | ~ spl66_58), inference(backward_demodulation, [], [f262, f834])).
fof(f262, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3444, plain, (~ spl66_42 | ~ spl66_58), inference(avatar_split_clause, [], [f3438, f832, f764])).
fof(f764, plain, (spl66_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_42])])).
fof(f3438, plain, (~ (e1 = op(e1, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f241, f834])).
fof(f241, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3426, plain, (~ spl66_31 | ~ spl66_63), inference(avatar_split_clause, [], [f3425, f853, f717])).
fof(f3425, plain, (~ (e2 = op(e2, e0)) | ~ spl66_63), inference(forward_demodulation, [], [f236, f855])).
fof(f855, plain, ((op(e0, e0) = e2) | ~ spl66_63), inference(avatar_component_clause, [], [f853])).
fof(f3424, plain, (~ spl66_32 | ~ spl66_63 | spl66_141), inference(avatar_split_clause, [], [f3344, f1476, f853, f721])).
fof(f1476, plain, (spl66_141 <=> (e3 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_141])])).
fof(f3344, plain, (~ (e3 = op(e2, e0)) | (~ spl66_63 | spl66_141)), inference(backward_demodulation, [], [f1478, f855])).
fof(f1478, plain, (~ (e3 = op(op(e0, e0), e0)) | spl66_141), inference(avatar_component_clause, [], [f1476])).
fof(f3419, plain, (~ spl66_3 | ~ spl66_11), inference(avatar_split_clause, [], [f3418, f632, f598])).
fof(f598, plain, (spl66_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_3])])).
fof(f3418, plain, (~ (e2 = op(e3, e3)) | ~ spl66_11), inference(forward_demodulation, [], [f281, f634])).
fof(f634, plain, ((e2 = op(e3, e1)) | ~ spl66_11), inference(avatar_component_clause, [], [f632])).
fof(f3410, plain, (~ spl66_8 | ~ spl66_16), inference(avatar_split_clause, [], [f3409, f653, f619])).
fof(f3409, plain, (~ (e3 = op(e3, e2)) | ~ spl66_16), inference(backward_demodulation, [], [f278, f655])).
fof(f278, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3405, plain, (~ spl66_2 | ~ spl66_24 | spl66_143), inference(avatar_contradiction_clause, [], [f3404])).
fof(f3404, plain, ($false | (~ spl66_2 | ~ spl66_24 | spl66_143)), inference(subsumption_resolution, [], [f3400, f596])).
fof(f596, plain, ((e1 = op(e3, e3)) | ~ spl66_2), inference(avatar_component_clause, [], [f594])).
fof(f3400, plain, (~ (e1 = op(e3, e3)) | (~ spl66_24 | spl66_143)), inference(backward_demodulation, [], [f1488, f689])).
fof(f689, plain, ((e3 = op(e2, e2)) | ~ spl66_24), inference(avatar_component_clause, [], [f687])).
fof(f3403, plain, (~ spl66_8 | ~ spl66_24), inference(avatar_split_clause, [], [f3399, f687, f619])).
fof(f3399, plain, (~ (e3 = op(e3, e2)) | ~ spl66_24), inference(backward_demodulation, [], [f252, f689])).
fof(f3398, plain, (~ spl66_21 | ~ spl66_25), inference(avatar_split_clause, [], [f3393, f692, f675])).
fof(f692, plain, (spl66_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_25])])).
fof(f3393, plain, (~ (e0 = op(e2, e2)) | ~ spl66_25), inference(backward_demodulation, [], [f274, f694])).
fof(f694, plain, ((e0 = op(e2, e1)) | ~ spl66_25), inference(avatar_component_clause, [], [f692])).
fof(f3385, plain, (~ spl66_38 | ~ spl66_42), inference(avatar_split_clause, [], [f3379, f764, f747])).
fof(f3379, plain, (~ (e1 = op(e1, e2)) | ~ spl66_42), inference(backward_demodulation, [], [f268, f766])).
fof(f766, plain, ((e1 = op(e1, e1)) | ~ spl66_42), inference(avatar_component_clause, [], [f764])).
fof(f3376, plain, (~ spl66_33 | ~ spl66_45), inference(avatar_split_clause, [], [f3372, f777, f726])).
fof(f777, plain, (spl66_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_45])])).
fof(f3372, plain, (~ (e0 = op(e1, e3)) | ~ spl66_45), inference(backward_demodulation, [], [f267, f779])).
fof(f779, plain, ((e0 = op(e1, e0)) | ~ spl66_45), inference(avatar_component_clause, [], [f777])).
fof(f267, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3368, plain, (~ spl66_33 | ~ spl66_49), inference(avatar_split_clause, [], [f3364, f794, f726])).
fof(f794, plain, (spl66_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_49])])).
fof(f3364, plain, (~ (e0 = op(e1, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f253, f796])).
fof(f796, plain, ((e0 = op(e0, e3)) | ~ spl66_49), inference(avatar_component_clause, [], [f794])).
fof(f253, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3356, plain, (~ spl66_52 | ~ spl66_60), inference(avatar_split_clause, [], [f3354, f840, f806])).
fof(f806, plain, (spl66_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_52])])).
fof(f840, plain, (spl66_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_60])])).
fof(f3354, plain, (~ (e3 = op(e0, e3)) | ~ spl66_60), inference(backward_demodulation, [], [f263, f842])).
fof(f842, plain, ((e3 = op(e0, e1)) | ~ spl66_60), inference(avatar_component_clause, [], [f840])).
fof(f263, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3355, plain, (~ spl66_28 | ~ spl66_60), inference(avatar_split_clause, [], [f3351, f840, f704])).
fof(f3351, plain, (~ (e3 = op(e2, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f242, f842])).
fof(f242, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3347, plain, (~ spl66_55 | ~ spl66_63), inference(avatar_split_clause, [], [f3340, f853, f819])).
fof(f819, plain, (spl66_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_55])])).
fof(f3340, plain, (~ (e2 = op(e0, e2)) | ~ spl66_63), inference(backward_demodulation, [], [f260, f855])).
fof(f3346, plain, (~ spl66_15 | ~ spl66_63), inference(avatar_split_clause, [], [f3338, f853, f649])).
fof(f649, plain, (spl66_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_15])])).
fof(f3338, plain, (~ (e2 = op(e3, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f237, f855])).
fof(f237, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3327, plain, (~ spl66_43 | ~ spl66_2 | spl66_151), inference(avatar_split_clause, [], [f3315, f1526, f594, f768])).
fof(f1526, plain, (spl66_151 <=> (e2 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_151])])).
fof(f3315, plain, (~ (e2 = op(e1, e1)) | (~ spl66_2 | spl66_151)), inference(backward_demodulation, [], [f1528, f596])).
fof(f1528, plain, (~ (e2 = op(op(e3, e3), op(e3, e3))) | spl66_151), inference(avatar_component_clause, [], [f1526])).
fof(f3308, plain, (~ spl66_1 | ~ spl66_9), inference(avatar_split_clause, [], [f3305, f624, f590])).
fof(f590, plain, (spl66_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_1])])).
fof(f624, plain, (spl66_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_9])])).
fof(f3305, plain, (~ (e0 = op(e3, e3)) | ~ spl66_9), inference(backward_demodulation, [], [f281, f626])).
fof(f626, plain, ((e0 = op(e3, e1)) | ~ spl66_9), inference(avatar_component_clause, [], [f624])).
fof(f3289, plain, (~ spl66_12 | ~ spl66_28), inference(avatar_split_clause, [], [f3286, f704, f636])).
fof(f3286, plain, (~ (e3 = op(e3, e1)) | ~ spl66_28), inference(backward_demodulation, [], [f246, f706])).
fof(f3285, plain, (~ spl66_26 | ~ spl66_30), inference(avatar_split_clause, [], [f3284, f713, f696])).
fof(f696, plain, (spl66_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_26])])).
fof(f713, plain, (spl66_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_30])])).
fof(f3284, plain, (~ (e1 = op(e2, e1)) | ~ spl66_30), inference(backward_demodulation, [], [f271, f715])).
fof(f715, plain, ((e1 = op(e2, e0)) | ~ spl66_30), inference(avatar_component_clause, [], [f713])).
fof(f3272, plain, (~ spl66_6 | ~ spl66_38), inference(avatar_split_clause, [], [f3268, f747, f611])).
fof(f611, plain, (spl66_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_6])])).
fof(f3268, plain, (~ (e1 = op(e3, e2)) | ~ spl66_38), inference(backward_demodulation, [], [f251, f749])).
fof(f251, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3262, plain, (~ spl66_51 | ~ spl66_55), inference(avatar_split_clause, [], [f3259, f819, f802])).
fof(f3259, plain, (~ (e2 = op(e0, e3)) | ~ spl66_55), inference(backward_demodulation, [], [f264, f821])).
fof(f821, plain, ((e2 = op(e0, e2)) | ~ spl66_55), inference(avatar_component_clause, [], [f819])).
fof(f264, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3261, plain, (~ spl66_23 | ~ spl66_55), inference(avatar_split_clause, [], [f3257, f819, f683])).
fof(f683, plain, (spl66_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_23])])).
fof(f3257, plain, (~ (e2 = op(e2, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f248, f821])).
fof(f3255, plain, (~ spl66_26 | ~ spl66_58), inference(avatar_split_clause, [], [f3249, f832, f696])).
fof(f3249, plain, (~ (e1 = op(e2, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f242, f834])).
fof(f3245, plain, (~ spl66_57 | ~ spl66_61), inference(avatar_split_clause, [], [f3235, f845, f828])).
fof(f845, plain, (spl66_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_61])])).
fof(f3235, plain, (~ (e0 = op(e0, e1)) | ~ spl66_61), inference(backward_demodulation, [], [f259, f847])).
fof(f847, plain, ((e0 = op(e0, e0)) | ~ spl66_61), inference(avatar_component_clause, [], [f845])).
fof(f259, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f3244, plain, (~ spl66_29 | ~ spl66_61), inference(avatar_split_clause, [], [f3234, f845, f709])).
fof(f709, plain, (spl66_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_29])])).
fof(f3234, plain, (~ (e0 = op(e2, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f236, f847])).
fof(f3224, plain, (~ spl66_35 | ~ spl66_43), inference(avatar_split_clause, [], [f3223, f768, f734])).
fof(f734, plain, (spl66_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_35])])).
fof(f3223, plain, (~ (e2 = op(e1, e3)) | ~ spl66_43), inference(forward_demodulation, [], [f269, f770])).
fof(f269, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3222, plain, (~ spl66_36 | ~ spl66_48), inference(avatar_split_clause, [], [f3221, f789, f738])).
fof(f789, plain, (spl66_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_48])])).
fof(f3221, plain, (~ (e3 = op(e1, e3)) | ~ spl66_48), inference(forward_demodulation, [], [f267, f791])).
fof(f791, plain, ((e3 = op(e1, e0)) | ~ spl66_48), inference(avatar_component_clause, [], [f789])).
fof(f3217, plain, (~ spl66_28 | ~ spl66_43 | spl66_132), inference(avatar_split_clause, [], [f3176, f1434, f768, f704])).
fof(f1434, plain, (spl66_132 <=> (e3 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_132])])).
fof(f3176, plain, (~ (e3 = op(e2, e1)) | (~ spl66_43 | spl66_132)), inference(backward_demodulation, [], [f1436, f770])).
fof(f1436, plain, (~ (e3 = op(op(e1, e1), e1)) | spl66_132), inference(avatar_component_clause, [], [f1434])).
fof(f3211, plain, (~ spl66_7 | ~ spl66_15), inference(avatar_split_clause, [], [f3209, f649, f615])).
fof(f615, plain, (spl66_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_7])])).
fof(f3209, plain, (~ (e2 = op(e3, e2)) | ~ spl66_15), inference(backward_demodulation, [], [f278, f651])).
fof(f651, plain, ((e2 = op(e3, e0)) | ~ spl66_15), inference(avatar_component_clause, [], [f649])).
fof(f3204, plain, (~ spl66_19 | ~ spl66_23), inference(avatar_split_clause, [], [f3195, f683, f666])).
fof(f3195, plain, (~ (e2 = op(e2, e3)) | ~ spl66_23), inference(backward_demodulation, [], [f276, f685])).
fof(f685, plain, ((e2 = op(e2, e2)) | ~ spl66_23), inference(avatar_component_clause, [], [f683])).
fof(f3203, plain, (~ spl66_7 | ~ spl66_23), inference(avatar_split_clause, [], [f3194, f683, f615])).
fof(f3194, plain, (~ (e2 = op(e3, e2)) | ~ spl66_23), inference(backward_demodulation, [], [f252, f685])).
fof(f3193, plain, (~ spl66_21 | ~ spl66_29), inference(avatar_split_clause, [], [f3188, f709, f675])).
fof(f3188, plain, (~ (e0 = op(e2, e2)) | ~ spl66_29), inference(backward_demodulation, [], [f272, f711])).
fof(f711, plain, ((e0 = op(e2, e0)) | ~ spl66_29), inference(avatar_component_clause, [], [f709])).
fof(f272, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3186, plain, (~ spl66_21 | ~ spl66_37), inference(avatar_split_clause, [], [f3181, f743, f675])).
fof(f743, plain, (spl66_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_37])])).
fof(f3181, plain, (~ (e0 = op(e2, e2)) | ~ spl66_37), inference(backward_demodulation, [], [f250, f745])).
fof(f745, plain, ((e0 = op(e1, e2)) | ~ spl66_37), inference(avatar_component_clause, [], [f743])).
fof(f3172, plain, (~ spl66_32 | ~ spl66_48), inference(avatar_split_clause, [], [f3168, f789, f721])).
fof(f3168, plain, (~ (e3 = op(e2, e0)) | ~ spl66_48), inference(backward_demodulation, [], [f238, f791])).
fof(f3157, plain, (~ spl66_41 | ~ spl66_57), inference(avatar_split_clause, [], [f3151, f828, f760])).
fof(f760, plain, (spl66_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_41])])).
fof(f3151, plain, (~ (e0 = op(e1, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f241, f830])).
fof(f3150, plain, (~ spl66_43 | ~ spl66_62 | spl66_147), inference(avatar_split_clause, [], [f3145, f1506, f849, f768])).
fof(f1506, plain, (spl66_147 <=> (e2 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_147])])).
fof(f3145, plain, (~ (e2 = op(e1, e1)) | (~ spl66_62 | spl66_147)), inference(backward_demodulation, [], [f1508, f851])).
fof(f1508, plain, (~ (e2 = op(op(e0, e0), op(e0, e0))) | spl66_147), inference(avatar_component_clause, [], [f1506])).
fof(f3148, plain, (~ spl66_54 | ~ spl66_62), inference(avatar_split_clause, [], [f3139, f849, f815])).
fof(f3139, plain, (~ (e1 = op(e0, e2)) | ~ spl66_62), inference(backward_demodulation, [], [f260, f851])).
fof(f3131, plain, (~ spl66_62 | ~ spl66_1 | spl66_145), inference(avatar_split_clause, [], [f3071, f1496, f590, f849])).
fof(f1496, plain, (spl66_145 <=> (e1 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_145])])).
fof(f3071, plain, (~ (op(e0, e0) = e1) | (~ spl66_1 | spl66_145)), inference(forward_demodulation, [], [f1498, f592])).
fof(f592, plain, ((e0 = op(e3, e3)) | ~ spl66_1), inference(avatar_component_clause, [], [f590])).
fof(f1498, plain, (~ (e1 = op(op(e3, e3), op(e3, e3))) | spl66_145), inference(avatar_component_clause, [], [f1496])).
fof(f3122, plain, (~ spl66_6 | ~ spl66_14), inference(avatar_split_clause, [], [f3119, f645, f611])).
fof(f3119, plain, (~ (e1 = op(e3, e2)) | ~ spl66_14), inference(backward_demodulation, [], [f278, f647])).
fof(f3118, plain, (~ spl66_21 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f3117])).
fof(f3117, plain, ($false | (~ spl66_21 | ~ spl66_24)), inference(subsumption_resolution, [], [f3116, f285])).
fof(f285, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f3116, plain, ((e0 = e3) | (~ spl66_21 | ~ spl66_24)), inference(backward_demodulation, [], [f689, f677])).
fof(f3108, plain, (~ spl66_37 | ~ spl66_41), inference(avatar_split_clause, [], [f3100, f760, f743])).
fof(f3100, plain, (~ (e0 = op(e1, e2)) | ~ spl66_41), inference(backward_demodulation, [], [f268, f762])).
fof(f762, plain, ((e0 = op(e1, e1)) | ~ spl66_41), inference(avatar_component_clause, [], [f760])).
fof(f3107, plain, (~ spl66_25 | ~ spl66_41), inference(avatar_split_clause, [], [f3099, f760, f692])).
fof(f3099, plain, (~ (e0 = op(e2, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f244, f762])).
fof(f3098, plain, (~ spl66_43 | ~ spl66_47), inference(avatar_split_clause, [], [f3094, f785, f768])).
fof(f3094, plain, (~ (e2 = op(e1, e1)) | ~ spl66_47), inference(backward_demodulation, [], [f265, f787])).
fof(f265, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3097, plain, (~ spl66_15 | ~ spl66_47), inference(avatar_split_clause, [], [f3093, f785, f649])).
fof(f3093, plain, (~ (e2 = op(e3, e0)) | ~ spl66_47), inference(backward_demodulation, [], [f239, f787])).
fof(f239, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3091, plain, (~ spl66_6 | ~ spl66_54), inference(avatar_split_clause, [], [f3088, f815, f611])).
fof(f3088, plain, (~ (e1 = op(e3, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f249, f817])).
fof(f3086, plain, (~ spl66_55 | ~ spl66_59), inference(avatar_split_clause, [], [f3083, f836, f819])).
fof(f836, plain, (spl66_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_59])])).
fof(f3083, plain, (~ (e2 = op(e0, e2)) | ~ spl66_59), inference(backward_demodulation, [], [f262, f838])).
fof(f838, plain, ((e2 = op(e0, e1)) | ~ spl66_59), inference(avatar_component_clause, [], [f836])).
fof(f3085, plain, (~ spl66_43 | ~ spl66_59), inference(avatar_split_clause, [], [f3081, f836, f768])).
fof(f3081, plain, (~ (e2 = op(e1, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f241, f838])).
fof(f3061, plain, (~ spl66_40 | ~ spl66_24), inference(avatar_split_clause, [], [f3060, f687, f755])).
fof(f3060, plain, (~ (e3 = op(e1, e2)) | ~ spl66_24), inference(forward_demodulation, [], [f250, f689])).
fof(f3058, plain, (~ spl66_32 | ~ spl66_24), inference(avatar_split_clause, [], [f3057, f687, f721])).
fof(f3057, plain, (~ (e3 = op(e2, e0)) | ~ spl66_24), inference(forward_demodulation, [], [f272, f689])).
fof(f3054, plain, (~ spl66_27 | ~ spl66_19), inference(avatar_split_clause, [], [f3053, f666, f700])).
fof(f3053, plain, (~ (e2 = op(e2, e1)) | ~ spl66_19), inference(forward_demodulation, [], [f275, f668])).
fof(f668, plain, ((e2 = op(e2, e3)) | ~ spl66_19), inference(avatar_component_clause, [], [f666])).
fof(f275, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3044, plain, (~ spl66_8 | ~ spl66_12), inference(avatar_split_clause, [], [f3043, f636, f619])).
fof(f3043, plain, (~ (e3 = op(e3, e2)) | ~ spl66_12), inference(backward_demodulation, [], [f280, f638])).
fof(f638, plain, ((e3 = op(e3, e1)) | ~ spl66_12), inference(avatar_component_clause, [], [f636])).
fof(f280, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3042, plain, (~ spl66_11 | ~ spl66_15), inference(avatar_split_clause, [], [f3039, f649, f632])).
fof(f3039, plain, (~ (e2 = op(e3, e1)) | ~ spl66_15), inference(backward_demodulation, [], [f277, f651])).
fof(f277, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3037, plain, (~ spl66_6 | ~ spl66_24 | spl66_136), inference(avatar_split_clause, [], [f3033, f1453, f687, f611])).
fof(f1453, plain, (spl66_136 <=> (e1 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_136])])).
fof(f3033, plain, (~ (e1 = op(e3, e2)) | (~ spl66_24 | spl66_136)), inference(backward_demodulation, [], [f1455, f689])).
fof(f1455, plain, (~ (e1 = op(op(e2, e2), e2)) | spl66_136), inference(avatar_component_clause, [], [f1453])).
fof(f3028, plain, (~ spl66_14 | ~ spl66_30), inference(avatar_split_clause, [], [f3025, f713, f645])).
fof(f3025, plain, (~ (e1 = op(e3, e0)) | ~ spl66_30), inference(backward_demodulation, [], [f240, f715])).
fof(f3027, plain, (~ spl66_18 | ~ spl66_30), inference(avatar_split_clause, [], [f3024, f713, f662])).
fof(f662, plain, (spl66_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_18])])).
fof(f3024, plain, (~ (e1 = op(e2, e3)) | ~ spl66_30), inference(backward_demodulation, [], [f273, f715])).
fof(f3026, plain, (~ spl66_22 | ~ spl66_30), inference(avatar_split_clause, [], [f3023, f713, f679])).
fof(f3023, plain, (~ (e1 = op(e2, e2)) | ~ spl66_30), inference(backward_demodulation, [], [f272, f715])).
fof(f2990, plain, (~ spl66_36 | ~ spl66_52), inference(avatar_split_clause, [], [f2988, f806, f738])).
fof(f2988, plain, (~ (e3 = op(e1, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f253, f808])).
fof(f808, plain, ((e3 = op(e0, e3)) | ~ spl66_52), inference(avatar_component_clause, [], [f806])).
fof(f2987, plain, (~ spl66_39 | ~ spl66_55), inference(avatar_split_clause, [], [f2982, f819, f751])).
fof(f2982, plain, (~ (e2 = op(e1, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f247, f821])).
fof(f2981, plain, (~ spl66_50 | ~ spl66_58), inference(avatar_split_clause, [], [f2977, f832, f798])).
fof(f798, plain, (spl66_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_50])])).
fof(f2977, plain, (~ (e1 = op(e0, e3)) | ~ spl66_58), inference(backward_demodulation, [], [f263, f834])).
fof(f2968, plain, (~ spl66_53 | ~ spl66_61), inference(avatar_split_clause, [], [f2959, f845, f811])).
fof(f811, plain, (spl66_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_53])])).
fof(f2959, plain, (~ (e0 = op(e0, e2)) | ~ spl66_61), inference(backward_demodulation, [], [f260, f847])).
fof(f2967, plain, (~ spl66_45 | ~ spl66_61), inference(avatar_split_clause, [], [f2955, f845, f777])).
fof(f2955, plain, (~ (e0 = op(e1, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f235, f847])).
fof(f235, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f2937, plain, (~ spl66_63 | ~ spl66_1 | spl66_151), inference(avatar_split_clause, [], [f2936, f1526, f590, f853])).
fof(f2936, plain, (~ (op(e0, e0) = e2) | (~ spl66_1 | spl66_151)), inference(forward_demodulation, [], [f1528, f592])).
fof(f2931, plain, (~ spl66_51 | ~ spl66_1 | spl66_138), inference(avatar_split_clause, [], [f2930, f1462, f590, f802])).
fof(f1462, plain, (spl66_138 <=> (e2 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_138])])).
fof(f2930, plain, (~ (e2 = op(e0, e3)) | (~ spl66_1 | spl66_138)), inference(forward_demodulation, [], [f1464, f592])).
fof(f1464, plain, (~ (e2 = op(op(e3, e3), e3)) | spl66_138), inference(avatar_component_clause, [], [f1462])).
fof(f2929, plain, (~ spl66_1 | ~ spl66_3), inference(avatar_contradiction_clause, [], [f2928])).
fof(f2928, plain, ($false | (~ spl66_1 | ~ spl66_3)), inference(subsumption_resolution, [], [f2927, f284])).
fof(f284, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f2927, plain, ((e0 = e2) | (~ spl66_1 | ~ spl66_3)), inference(backward_demodulation, [], [f600, f592])).
fof(f600, plain, ((e2 = op(e3, e3)) | ~ spl66_3), inference(avatar_component_clause, [], [f598])).
fof(f2921, plain, (~ spl66_17 | ~ spl66_25), inference(avatar_split_clause, [], [f2918, f692, f658])).
fof(f2918, plain, (~ (e0 = op(e2, e3)) | ~ spl66_25), inference(backward_demodulation, [], [f275, f694])).
fof(f2915, plain, (~ spl66_31 | ~ spl66_32), inference(avatar_contradiction_clause, [], [f2914])).
fof(f2914, plain, ($false | (~ spl66_31 | ~ spl66_32)), inference(subsumption_resolution, [], [f2913, f288])).
fof(f288, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f2913, plain, ((e2 = e3) | (~ spl66_31 | ~ spl66_32)), inference(backward_demodulation, [], [f723, f719])).
fof(f2904, plain, (~ spl66_34 | ~ spl66_42), inference(avatar_split_clause, [], [f2897, f764, f730])).
fof(f2897, plain, (~ (e1 = op(e1, e3)) | ~ spl66_42), inference(backward_demodulation, [], [f269, f766])).
fof(f2893, plain, (~ spl66_37 | ~ spl66_45), inference(avatar_split_clause, [], [f2887, f777, f743])).
fof(f2887, plain, (~ (e0 = op(e1, e2)) | ~ spl66_45), inference(backward_demodulation, [], [f266, f779])).
fof(f2892, plain, (~ spl66_41 | ~ spl66_45), inference(avatar_split_clause, [], [f2886, f777, f760])).
fof(f2886, plain, (~ (e0 = op(e1, e1)) | ~ spl66_45), inference(backward_demodulation, [], [f265, f779])).
fof(f2885, plain, (~ spl66_18 | ~ spl66_50), inference(avatar_split_clause, [], [f2881, f798, f662])).
fof(f2881, plain, (~ (e1 = op(e2, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f254, f800])).
fof(f800, plain, ((e1 = op(e0, e3)) | ~ spl66_50), inference(avatar_component_clause, [], [f798])).
fof(f2884, plain, (~ spl66_34 | ~ spl66_50), inference(avatar_split_clause, [], [f2880, f798, f730])).
fof(f2880, plain, (~ (e1 = op(e1, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f253, f800])).
fof(f2879, plain, (~ spl66_5 | ~ spl66_53), inference(avatar_split_clause, [], [f2873, f811, f607])).
fof(f2873, plain, (~ (e0 = op(e3, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f249, f813])).
fof(f813, plain, ((e0 = op(e0, e2)) | ~ spl66_53), inference(avatar_component_clause, [], [f811])).
fof(f2878, plain, (~ spl66_37 | ~ spl66_53), inference(avatar_split_clause, [], [f2872, f811, f743])).
fof(f2872, plain, (~ (e0 = op(e1, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f247, f813])).
fof(f2870, plain, (~ spl66_12 | ~ spl66_60), inference(avatar_split_clause, [], [f2867, f840, f636])).
fof(f2867, plain, (~ (e3 = op(e3, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f243, f842])).
fof(f243, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f2863, plain, (~ spl66_22 | ~ spl66_63 | spl66_140), inference(avatar_contradiction_clause, [], [f2862])).
fof(f2862, plain, ($false | (~ spl66_22 | ~ spl66_63 | spl66_140)), inference(subsumption_resolution, [], [f2857, f681])).
fof(f681, plain, ((e1 = op(e2, e2)) | ~ spl66_22), inference(avatar_component_clause, [], [f679])).
fof(f2857, plain, (~ (e1 = op(e2, e2)) | (~ spl66_63 | spl66_140)), inference(backward_demodulation, [], [f1474, f855])).
fof(f1474, plain, (~ (e1 = op(op(e0, e0), op(e0, e0))) | spl66_140), inference(avatar_component_clause, [], [f1472])).
fof(f1472, plain, (spl66_140 <=> (e1 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_140])])).
fof(f2860, plain, (~ spl66_47 | ~ spl66_63), inference(avatar_split_clause, [], [f2852, f853, f785])).
fof(f2852, plain, (~ (e2 = op(e1, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f235, f855])).
fof(f2851, plain, (~ spl66_40 | ~ spl66_22 | spl66_135), inference(avatar_split_clause, [], [f2850, f1448, f679, f755])).
fof(f2850, plain, (~ (e3 = op(e1, e2)) | (~ spl66_22 | spl66_135)), inference(forward_demodulation, [], [f1450, f681])).
fof(f2830, plain, (~ spl66_35 | ~ spl66_3), inference(avatar_split_clause, [], [f2829, f598, f734])).
fof(f2829, plain, (~ (e2 = op(e1, e3)) | ~ spl66_3), inference(forward_demodulation, [], [f257, f600])).
fof(f2827, plain, (~ spl66_18 | ~ spl66_22), inference(avatar_split_clause, [], [f2805, f679, f662])).
fof(f2805, plain, (~ (e1 = op(e2, e3)) | ~ spl66_22), inference(backward_demodulation, [], [f276, f681])).
fof(f2820, plain, (~ spl66_6 | ~ spl66_22), inference(avatar_split_clause, [], [f2819, f679, f611])).
fof(f2819, plain, (~ (e1 = op(e3, e2)) | ~ spl66_22), inference(forward_demodulation, [], [f252, f681])).
fof(f2813, plain, (~ spl66_22 | ~ spl66_41 | spl66_134), inference(avatar_contradiction_clause, [], [f2812])).
fof(f2812, plain, ($false | (~ spl66_22 | ~ spl66_41 | spl66_134)), inference(subsumption_resolution, [], [f2808, f762])).
fof(f2808, plain, (~ (e0 = op(e1, e1)) | (~ spl66_22 | spl66_134)), inference(backward_demodulation, [], [f1446, f681])).
fof(f1446, plain, (~ (e0 = op(op(e2, e2), op(e2, e2))) | spl66_134), inference(avatar_component_clause, [], [f1444])).
fof(f1444, plain, (spl66_134 <=> (e0 = op(op(e2, e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_134])])).
fof(f2804, plain, (~ spl66_23 | ~ spl66_27), inference(avatar_split_clause, [], [f2802, f700, f683])).
fof(f2802, plain, (~ (e2 = op(e2, e2)) | ~ spl66_27), inference(backward_demodulation, [], [f274, f702])).
fof(f2780, plain, (~ spl66_49 | ~ spl66_61), inference(avatar_split_clause, [], [f2772, f845, f794])).
fof(f2772, plain, (~ (e0 = op(e0, e3)) | ~ spl66_61), inference(backward_demodulation, [], [f261, f847])).
fof(f2768, plain, (~ spl66_59 | ~ spl66_41 | spl66_133), inference(avatar_split_clause, [], [f2767, f1439, f760, f836])).
fof(f2767, plain, (~ (e2 = op(e0, e1)) | (~ spl66_41 | spl66_133)), inference(forward_demodulation, [], [f1441, f762])).
fof(f2735, plain, (~ spl66_2 | ~ spl66_14), inference(avatar_split_clause, [], [f2733, f645, f594])).
fof(f2733, plain, (~ (e1 = op(e3, e3)) | ~ spl66_14), inference(backward_demodulation, [], [f279, f647])).
fof(f2723, plain, (~ spl66_35 | ~ spl66_47), inference(avatar_split_clause, [], [f2721, f785, f734])).
fof(f2721, plain, (~ (e2 = op(e1, e3)) | ~ spl66_47), inference(backward_demodulation, [], [f267, f787])).
fof(f2717, plain, (~ spl66_11 | ~ spl66_59), inference(avatar_split_clause, [], [f2716, f836, f632])).
fof(f2716, plain, (~ (e2 = op(e3, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f243, f838])).
fof(f2711, plain, (~ spl66_3 | ~ spl66_64 | spl66_147), inference(avatar_split_clause, [], [f2705, f1506, f857, f598])).
fof(f2705, plain, (~ (e2 = op(e3, e3)) | (~ spl66_64 | spl66_147)), inference(backward_demodulation, [], [f1508, f859])).
fof(f2710, plain, (~ spl66_2 | ~ spl66_64 | spl66_140), inference(avatar_split_clause, [], [f2704, f1472, f857, f594])).
fof(f2704, plain, (~ (e1 = op(e3, e3)) | (~ spl66_64 | spl66_140)), inference(backward_demodulation, [], [f1474, f859])).
fof(f2708, plain, (~ spl66_16 | ~ spl66_64), inference(avatar_split_clause, [], [f2701, f857, f653])).
fof(f2701, plain, (~ (e3 = op(e3, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f237, f859])).
fof(f2691, plain, (~ spl66_64 | ~ spl66_41 | spl66_153), inference(avatar_split_clause, [], [f2606, f1538, f760, f857])).
fof(f1538, plain, (spl66_153 <=> (e3 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_153])])).
fof(f2606, plain, (~ (op(e0, e0) = e3) | (~ spl66_41 | spl66_153)), inference(backward_demodulation, [], [f1540, f762])).
fof(f1540, plain, (~ (e3 = op(op(e1, e1), op(e1, e1))) | spl66_153), inference(avatar_component_clause, [], [f1538])).
fof(f2680, plain, (~ spl66_4 | ~ spl66_20), inference(avatar_split_clause, [], [f2679, f670, f602])).
fof(f670, plain, (spl66_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_20])])).
fof(f2679, plain, (~ (e3 = op(e3, e3)) | ~ spl66_20), inference(forward_demodulation, [], [f258, f672])).
fof(f672, plain, ((e3 = op(e2, e3)) | ~ spl66_20), inference(avatar_component_clause, [], [f670])).
fof(f258, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f2677, plain, (~ spl66_22 | ~ spl66_23), inference(avatar_contradiction_clause, [], [f2676])).
fof(f2676, plain, ($false | (~ spl66_22 | ~ spl66_23)), inference(subsumption_resolution, [], [f2675, f286])).
fof(f286, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f2675, plain, ((e1 = e2) | (~ spl66_22 | ~ spl66_23)), inference(backward_demodulation, [], [f685, f681])).
fof(f2668, plain, (~ spl66_60 | ~ spl66_41 | spl66_132), inference(avatar_split_clause, [], [f2667, f1434, f760, f840])).
fof(f2667, plain, (~ (e3 = op(e0, e1)) | (~ spl66_41 | spl66_132)), inference(forward_demodulation, [], [f1436, f762])).
fof(f2665, plain, (~ spl66_2 | ~ spl66_35 | spl66_138), inference(avatar_contradiction_clause, [], [f2664])).
fof(f2664, plain, ($false | (~ spl66_2 | ~ spl66_35 | spl66_138)), inference(subsumption_resolution, [], [f2663, f736])).
fof(f736, plain, ((e2 = op(e1, e3)) | ~ spl66_35), inference(avatar_component_clause, [], [f734])).
fof(f2663, plain, (~ (e2 = op(e1, e3)) | (~ spl66_2 | spl66_138)), inference(forward_demodulation, [], [f1464, f596])).
fof(f2633, plain, (~ spl66_22 | ~ spl66_26), inference(avatar_split_clause, [], [f2632, f696, f679])).
fof(f2632, plain, (~ (e1 = op(e2, e2)) | ~ spl66_26), inference(backward_demodulation, [], [f274, f698])).
fof(f698, plain, ((e1 = op(e2, e1)) | ~ spl66_26), inference(avatar_component_clause, [], [f696])).
fof(f2628, plain, (~ spl66_25 | ~ spl66_29), inference(avatar_split_clause, [], [f2621, f709, f692])).
fof(f2621, plain, (~ (e0 = op(e2, e1)) | ~ spl66_29), inference(backward_demodulation, [], [f271, f711])).
fof(f2627, plain, (~ spl66_13 | ~ spl66_29), inference(avatar_split_clause, [], [f2620, f709, f641])).
fof(f641, plain, (spl66_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_13])])).
fof(f2620, plain, (~ (e0 = op(e3, e0)) | ~ spl66_29), inference(backward_demodulation, [], [f240, f711])).
fof(f2619, plain, (~ spl66_19 | ~ spl66_35), inference(avatar_split_clause, [], [f2618, f734, f666])).
fof(f2618, plain, (~ (e2 = op(e2, e3)) | ~ spl66_35), inference(backward_demodulation, [], [f256, f736])).
fof(f2616, plain, (~ spl66_36 | ~ spl66_40), inference(avatar_split_clause, [], [f2614, f755, f738])).
fof(f2614, plain, (~ (e3 = op(e1, e3)) | ~ spl66_40), inference(backward_demodulation, [], [f270, f757])).
fof(f757, plain, ((e3 = op(e1, e2)) | ~ spl66_40), inference(avatar_component_clause, [], [f755])).
fof(f2610, plain, (~ spl66_41 | ~ spl66_63 | spl66_149), inference(avatar_contradiction_clause, [], [f2609])).
fof(f2609, plain, ($false | (~ spl66_41 | ~ spl66_63 | spl66_149)), inference(subsumption_resolution, [], [f2604, f855])).
fof(f2604, plain, (~ (op(e0, e0) = e2) | (~ spl66_41 | spl66_149)), inference(backward_demodulation, [], [f1518, f762])).
fof(f1518, plain, (~ (e2 = op(op(e1, e1), op(e1, e1))) | spl66_149), inference(avatar_component_clause, [], [f1516])).
fof(f1516, plain, (spl66_149 <=> (e2 = op(op(e1, e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_149])])).
fof(f2595, plain, (~ spl66_42 | ~ spl66_46), inference(avatar_split_clause, [], [f2590, f781, f764])).
fof(f2590, plain, (~ (e1 = op(e1, e1)) | ~ spl66_46), inference(backward_demodulation, [], [f265, f783])).
fof(f2587, plain, (~ spl66_22 | ~ spl66_54), inference(avatar_split_clause, [], [f2584, f815, f679])).
fof(f2584, plain, (~ (e1 = op(e2, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f248, f817])).
fof(f2565, plain, (~ spl66_64 | ~ spl66_60), inference(avatar_split_clause, [], [f2564, f840, f857])).
fof(f2564, plain, (~ (op(e0, e0) = e3) | ~ spl66_60), inference(forward_demodulation, [], [f259, f842])).
fof(f2561, plain, (~ spl66_41 | ~ spl66_2 | spl66_137), inference(avatar_split_clause, [], [f2422, f1458, f594, f760])).
fof(f1458, plain, (spl66_137 <=> (e0 = op(op(e3, e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_137])])).
fof(f2422, plain, (~ (e0 = op(e1, e1)) | (~ spl66_2 | spl66_137)), inference(backward_demodulation, [], [f1460, f596])).
fof(f1460, plain, (~ (e0 = op(op(e3, e3), op(e3, e3))) | spl66_137), inference(avatar_component_clause, [], [f1458])).
fof(f2531, plain, (~ spl66_62 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f2530])).
fof(f2530, plain, ($false | (~ spl66_62 | ~ spl66_64)), inference(subsumption_resolution, [], [f2529, f287])).
fof(f2529, plain, ((e1 = e3) | (~ spl66_62 | ~ spl66_64)), inference(backward_demodulation, [], [f859, f851])).
fof(f2480, plain, (~ spl66_9 | ~ spl66_13), inference(avatar_split_clause, [], [f2477, f641, f624])).
fof(f2477, plain, (~ (e0 = op(e3, e1)) | ~ spl66_13), inference(backward_demodulation, [], [f277, f643])).
fof(f643, plain, ((e0 = op(e3, e0)) | ~ spl66_13), inference(avatar_component_clause, [], [f641])).
fof(f2463, plain, (~ spl66_5 | ~ spl66_37), inference(avatar_split_clause, [], [f2460, f743, f607])).
fof(f2460, plain, (~ (e0 = op(e3, e2)) | ~ spl66_37), inference(backward_demodulation, [], [f251, f745])).
fof(f2458, plain, (~ spl66_42 | ~ spl66_44), inference(avatar_contradiction_clause, [], [f2457])).
fof(f2457, plain, ($false | (~ spl66_42 | ~ spl66_44)), inference(subsumption_resolution, [], [f2456, f287])).
fof(f2456, plain, ((e1 = e3) | (~ spl66_42 | ~ spl66_44)), inference(backward_demodulation, [], [f774, f766])).
fof(f2454, plain, (~ spl66_46 | ~ spl66_47), inference(avatar_contradiction_clause, [], [f2453])).
fof(f2453, plain, ($false | (~ spl66_46 | ~ spl66_47)), inference(subsumption_resolution, [], [f2452, f286])).
fof(f2452, plain, ((e1 = e2) | (~ spl66_46 | ~ spl66_47)), inference(forward_demodulation, [], [f787, f783])).
fof(f2443, plain, (~ spl66_30 | ~ spl66_46), inference(avatar_split_clause, [], [f2442, f781, f713])).
fof(f2442, plain, (~ (e1 = op(e2, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f238, f783])).
fof(f2419, plain, (~ spl66_1 | ~ spl66_5), inference(avatar_split_clause, [], [f2417, f607, f590])).
fof(f2417, plain, (~ (e0 = op(e3, e3)) | ~ spl66_5), inference(backward_demodulation, [], [f282, f609])).
fof(f609, plain, ((e0 = op(e3, e2)) | ~ spl66_5), inference(avatar_component_clause, [], [f607])).
fof(f2413, plain, (~ spl66_5 | ~ spl66_9), inference(avatar_split_clause, [], [f2408, f624, f607])).
fof(f2408, plain, (~ (e0 = op(e3, e2)) | ~ spl66_9), inference(backward_demodulation, [], [f280, f626])).
fof(f2407, plain, (~ spl66_15 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f2406])).
fof(f2406, plain, ($false | (~ spl66_15 | ~ spl66_16)), inference(subsumption_resolution, [], [f2405, f288])).
fof(f2405, plain, ((e2 = e3) | (~ spl66_15 | ~ spl66_16)), inference(backward_demodulation, [], [f655, f651])).
fof(f2404, plain, (~ spl66_12 | ~ spl66_16), inference(avatar_split_clause, [], [f2403, f653, f636])).
fof(f2403, plain, (~ (e3 = op(e3, e1)) | ~ spl66_16), inference(backward_demodulation, [], [f277, f655])).
fof(f2395, plain, (~ spl66_9 | ~ spl66_44 | spl66_150), inference(avatar_split_clause, [], [f2388, f1521, f772, f624])).
fof(f1521, plain, (spl66_150 <=> (e0 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_150])])).
fof(f2388, plain, (~ (e0 = op(e3, e1)) | (~ spl66_44 | spl66_150)), inference(backward_demodulation, [], [f1523, f774])).
fof(f1523, plain, (~ (e0 = op(op(e1, e1), e1)) | spl66_150), inference(avatar_component_clause, [], [f1521])).
fof(f2391, plain, (~ spl66_1 | ~ spl66_44 | spl66_131), inference(avatar_split_clause, [], [f2384, f1430, f772, f590])).
fof(f2384, plain, (~ (e0 = op(e3, e3)) | (~ spl66_44 | spl66_131)), inference(backward_demodulation, [], [f1432, f774])).
fof(f2381, plain, (~ spl66_45 | ~ spl66_46), inference(avatar_contradiction_clause, [], [f2380])).
fof(f2380, plain, ($false | (~ spl66_45 | ~ spl66_46)), inference(subsumption_resolution, [], [f2379, f283])).
fof(f2379, plain, ((e0 = e1) | (~ spl66_45 | ~ spl66_46)), inference(backward_demodulation, [], [f783, f779])).
fof(f2378, plain, (~ spl66_14 | ~ spl66_46), inference(avatar_split_clause, [], [f2375, f781, f645])).
fof(f2375, plain, (~ (e1 = op(e3, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f239, f783])).
fof(f2365, plain, (~ spl66_15 | ~ spl66_64 | spl66_142), inference(avatar_split_clause, [], [f2364, f1481, f857, f649])).
fof(f2364, plain, (~ (e2 = op(e3, e0)) | (~ spl66_64 | spl66_142)), inference(forward_demodulation, [], [f1483, f859])).
fof(f2353, plain, (~ spl66_39 | ~ spl66_23), inference(avatar_split_clause, [], [f2352, f683, f751])).
fof(f2352, plain, (~ (e2 = op(e1, e2)) | ~ spl66_23), inference(forward_demodulation, [], [f250, f685])).
fof(f2307, plain, (~ spl66_37 | ~ spl66_38), inference(avatar_contradiction_clause, [], [f2306])).
fof(f2306, plain, ($false | (~ spl66_37 | ~ spl66_38)), inference(subsumption_resolution, [], [f2305, f283])).
fof(f2305, plain, ((e0 = e1) | (~ spl66_37 | ~ spl66_38)), inference(backward_demodulation, [], [f749, f745])).
fof(f2304, plain, (~ spl66_38 | ~ spl66_40), inference(avatar_contradiction_clause, [], [f2303])).
fof(f2303, plain, ($false | (~ spl66_38 | ~ spl66_40)), inference(subsumption_resolution, [], [f2302, f287])).
fof(f2302, plain, ((e1 = e3) | (~ spl66_38 | ~ spl66_40)), inference(forward_demodulation, [], [f757, f749])).
fof(f2296, plain, (~ spl66_49 | ~ spl66_51), inference(avatar_contradiction_clause, [], [f2295])).
fof(f2295, plain, ($false | (~ spl66_49 | ~ spl66_51)), inference(subsumption_resolution, [], [f2294, f284])).
fof(f2294, plain, ((e0 = e2) | (~ spl66_49 | ~ spl66_51)), inference(forward_demodulation, [], [f804, f796])).
fof(f2289, plain, (~ spl66_53 | ~ spl66_57), inference(avatar_split_clause, [], [f2285, f828, f811])).
fof(f2285, plain, (~ (e0 = op(e0, e2)) | ~ spl66_57), inference(backward_demodulation, [], [f262, f830])).
fof(f2283, plain, (~ spl66_48 | ~ spl66_64), inference(avatar_split_clause, [], [f2280, f857, f789])).
fof(f2280, plain, (~ (e3 = op(e1, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f235, f859])).
fof(f2277, plain, (~ spl66_57 | ~ spl66_49), inference(avatar_split_clause, [], [f2276, f794, f828])).
fof(f2276, plain, (~ (e0 = op(e0, e1)) | ~ spl66_49), inference(forward_demodulation, [], [f263, f796])).
fof(f2271, plain, (~ spl66_53 | ~ spl66_49), inference(avatar_split_clause, [], [f2270, f794, f811])).
fof(f2270, plain, (~ (e0 = op(e0, e2)) | ~ spl66_49), inference(forward_demodulation, [], [f264, f796])).
fof(f2266, plain, (~ spl66_1 | ~ spl66_49), inference(avatar_split_clause, [], [f2265, f794, f590])).
fof(f2265, plain, (~ (e0 = op(e3, e3)) | ~ spl66_49), inference(forward_demodulation, [], [f255, f796])).
fof(f255, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f2258, plain, (~ spl66_31 | ~ spl66_23), inference(avatar_split_clause, [], [f2257, f683, f717])).
fof(f2257, plain, (~ (e2 = op(e2, e0)) | ~ spl66_23), inference(forward_demodulation, [], [f272, f685])).
fof(f2254, plain, (~ spl66_18 | ~ spl66_26), inference(avatar_split_clause, [], [f2253, f696, f662])).
fof(f2253, plain, (~ (e1 = op(e2, e3)) | ~ spl66_26), inference(forward_demodulation, [], [f275, f698])).
fof(f2244, plain, (~ spl66_14 | ~ spl66_64 | spl66_148), inference(avatar_split_clause, [], [f2195, f1511, f857, f645])).
fof(f1511, plain, (spl66_148 <=> (e1 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_148])])).
fof(f2195, plain, (~ (e1 = op(e3, e0)) | (~ spl66_64 | spl66_148)), inference(backward_demodulation, [], [f1513, f859])).
fof(f1513, plain, (~ (e1 = op(op(e0, e0), e0)) | spl66_148), inference(avatar_component_clause, [], [f1511])).
fof(f2233, plain, (~ spl66_2 | ~ spl66_6), inference(avatar_split_clause, [], [f2232, f611, f594])).
fof(f2232, plain, (~ (e1 = op(e3, e3)) | ~ spl66_6), inference(backward_demodulation, [], [f282, f613])).
fof(f613, plain, ((e1 = op(e3, e2)) | ~ spl66_6), inference(avatar_component_clause, [], [f611])).
fof(f2208, plain, (~ spl66_49 | ~ spl66_50), inference(avatar_contradiction_clause, [], [f2207])).
fof(f2207, plain, ($false | (~ spl66_49 | ~ spl66_50)), inference(subsumption_resolution, [], [f2206, f283])).
fof(f2206, plain, ((e0 = e1) | (~ spl66_49 | ~ spl66_50)), inference(backward_demodulation, [], [f800, f796])).
fof(f2200, plain, (~ spl66_63 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f2199])).
fof(f2199, plain, ($false | (~ spl66_63 | ~ spl66_64)), inference(subsumption_resolution, [], [f2198, f288])).
fof(f2198, plain, ((e2 = e3) | (~ spl66_63 | ~ spl66_64)), inference(backward_demodulation, [], [f859, f855])).
fof(f2187, plain, (~ spl66_63 | ~ spl66_59), inference(avatar_split_clause, [], [f2186, f836, f853])).
fof(f2186, plain, (~ (op(e0, e0) = e2) | ~ spl66_59), inference(forward_demodulation, [], [f259, f838])).
fof(f2184, plain, (~ spl66_62 | ~ spl66_50), inference(avatar_split_clause, [], [f2183, f798, f849])).
fof(f2183, plain, (~ (op(e0, e0) = e1) | ~ spl66_50), inference(forward_demodulation, [], [f261, f800])).
fof(f2175, plain, (~ spl66_16 | ~ spl66_48), inference(avatar_split_clause, [], [f2174, f789, f653])).
fof(f2174, plain, (~ (e3 = op(e3, e0)) | ~ spl66_48), inference(forward_demodulation, [], [f239, f791])).
fof(f2162, plain, (~ spl66_2 | ~ spl66_50), inference(avatar_split_clause, [], [f2161, f798, f594])).
fof(f2161, plain, (~ (e1 = op(e3, e3)) | ~ spl66_50), inference(forward_demodulation, [], [f255, f800])).
fof(f2137, plain, (~ spl66_44 | ~ spl66_48), inference(avatar_split_clause, [], [f2136, f789, f772])).
fof(f2136, plain, (~ (e3 = op(e1, e1)) | ~ spl66_48), inference(backward_demodulation, [], [f265, f791])).
fof(f2044, plain, (~ spl66_10 | ~ spl66_2), inference(avatar_split_clause, [], [f2043, f594, f628])).
fof(f628, plain, (spl66_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_10])])).
fof(f2043, plain, (~ (e1 = op(e3, e1)) | ~ spl66_2), inference(forward_demodulation, [], [f281, f596])).
fof(f2042, plain, (~ spl66_11 | ~ spl66_7), inference(avatar_split_clause, [], [f1998, f615, f632])).
fof(f1998, plain, (~ (e2 = op(e3, e1)) | ~ spl66_7), inference(forward_demodulation, [], [f280, f617])).
fof(f617, plain, ((e2 = op(e3, e2)) | ~ spl66_7), inference(avatar_component_clause, [], [f615])).
fof(f2034, plain, (~ spl66_2 | ~ spl66_33 | spl66_146), inference(avatar_contradiction_clause, [], [f2033])).
fof(f2033, plain, ($false | (~ spl66_2 | ~ spl66_33 | spl66_146)), inference(subsumption_resolution, [], [f2027, f728])).
fof(f2027, plain, (~ (e0 = op(e1, e3)) | (~ spl66_2 | spl66_146)), inference(backward_demodulation, [], [f1503, f596])).
fof(f1503, plain, (~ (e0 = op(op(e3, e3), e3)) | spl66_146), inference(avatar_component_clause, [], [f1501])).
fof(f1501, plain, (spl66_146 <=> (e0 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_146])])).
fof(f2020, plain, (~ spl66_51 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f2019])).
fof(f2019, plain, ($false | (~ spl66_51 | ~ spl66_52)), inference(subsumption_resolution, [], [f2018, f288])).
fof(f2018, plain, ((e2 = e3) | (~ spl66_51 | ~ spl66_52)), inference(forward_demodulation, [], [f808, f804])).
fof(f2003, plain, (~ spl66_55 | ~ spl66_7), inference(avatar_split_clause, [], [f2002, f615, f819])).
fof(f2002, plain, (~ (e2 = op(e0, e2)) | ~ spl66_7), inference(forward_demodulation, [], [f249, f617])).
fof(f1984, plain, (~ spl66_7 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f1983])).
fof(f1983, plain, ($false | (~ spl66_7 | ~ spl66_8)), inference(subsumption_resolution, [], [f1982, f288])).
fof(f1982, plain, ((e2 = e3) | (~ spl66_7 | ~ spl66_8)), inference(backward_demodulation, [], [f621, f617])).
fof(f621, plain, ((e3 = op(e3, e2)) | ~ spl66_8), inference(avatar_component_clause, [], [f619])).
fof(f1979, plain, (~ spl66_18 | ~ spl66_20), inference(avatar_contradiction_clause, [], [f1978])).
fof(f1978, plain, ($false | (~ spl66_18 | ~ spl66_20)), inference(subsumption_resolution, [], [f1976, f287])).
fof(f1976, plain, ((e1 = e3) | (~ spl66_18 | ~ spl66_20)), inference(backward_demodulation, [], [f672, f664])).
fof(f664, plain, ((e1 = op(e2, e3)) | ~ spl66_18), inference(avatar_component_clause, [], [f662])).
fof(f1968, plain, (~ spl66_61 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f1967])).
fof(f1967, plain, ($false | (~ spl66_61 | ~ spl66_64)), inference(subsumption_resolution, [], [f1966, f285])).
fof(f1966, plain, ((e0 = e3) | (~ spl66_61 | ~ spl66_64)), inference(forward_demodulation, [], [f859, f847])).
fof(f1955, plain, (~ spl66_53 | ~ spl66_21), inference(avatar_split_clause, [], [f1863, f675, f811])).
fof(f1863, plain, (~ (e0 = op(e0, e2)) | ~ spl66_21), inference(forward_demodulation, [], [f248, f677])).
fof(f1929, plain, (~ spl66_50 | ~ spl66_1 | spl66_139), inference(avatar_split_clause, [], [f1928, f1467, f590, f798])).
fof(f1467, plain, (spl66_139 <=> (e1 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_139])])).
fof(f1928, plain, (~ (e1 = op(e0, e3)) | (~ spl66_1 | spl66_139)), inference(forward_demodulation, [], [f1469, f592])).
fof(f1469, plain, (~ (e1 = op(op(e3, e3), e3)) | spl66_139), inference(avatar_component_clause, [], [f1467])).
fof(f1918, plain, (~ spl66_1 | ~ spl66_2), inference(avatar_contradiction_clause, [], [f1917])).
fof(f1917, plain, ($false | (~ spl66_1 | ~ spl66_2)), inference(subsumption_resolution, [], [f1916, f283])).
fof(f1916, plain, ((e0 = e1) | (~ spl66_1 | ~ spl66_2)), inference(backward_demodulation, [], [f596, f592])).
fof(f1915, plain, (~ spl66_2 | ~ spl66_3), inference(avatar_contradiction_clause, [], [f1914])).
fof(f1914, plain, ($false | (~ spl66_2 | ~ spl66_3)), inference(subsumption_resolution, [], [f1913, f286])).
fof(f1913, plain, ((e1 = e2) | (~ spl66_2 | ~ spl66_3)), inference(backward_demodulation, [], [f600, f596])).
fof(f1900, plain, (~ spl66_3 | ~ spl66_44 | spl66_149), inference(avatar_contradiction_clause, [], [f1899])).
fof(f1899, plain, ($false | (~ spl66_3 | ~ spl66_44 | spl66_149)), inference(subsumption_resolution, [], [f1896, f600])).
fof(f1896, plain, (~ (e2 = op(e3, e3)) | (~ spl66_44 | spl66_149)), inference(backward_demodulation, [], [f1518, f774])).
fof(f1882, plain, (~ spl66_55 | ~ spl66_56), inference(avatar_contradiction_clause, [], [f1881])).
fof(f1881, plain, ($false | (~ spl66_55 | ~ spl66_56)), inference(subsumption_resolution, [], [f1880, f288])).
fof(f1880, plain, ((e2 = e3) | (~ spl66_55 | ~ spl66_56)), inference(backward_demodulation, [], [f825, f821])).
fof(f1875, plain, (~ spl66_10 | ~ spl66_58), inference(avatar_split_clause, [], [f1869, f832, f628])).
fof(f1869, plain, (~ (e1 = op(e3, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f243, f834])).
fof(f1843, plain, (~ spl66_5 | ~ spl66_8), inference(avatar_contradiction_clause, [], [f1842])).
fof(f1842, plain, ($false | (~ spl66_5 | ~ spl66_8)), inference(subsumption_resolution, [], [f1841, f285])).
fof(f1841, plain, ((e0 = e3) | (~ spl66_5 | ~ spl66_8)), inference(backward_demodulation, [], [f621, f609])).
fof(f1837, plain, (~ spl66_25 | ~ spl66_27), inference(avatar_contradiction_clause, [], [f1836])).
fof(f1836, plain, ($false | (~ spl66_25 | ~ spl66_27)), inference(subsumption_resolution, [], [f1835, f284])).
fof(f1835, plain, ((e0 = e2) | (~ spl66_25 | ~ spl66_27)), inference(backward_demodulation, [], [f702, f694])).
fof(f1815, plain, (~ spl66_61 | ~ spl66_63), inference(avatar_contradiction_clause, [], [f1814])).
fof(f1814, plain, ($false | (~ spl66_61 | ~ spl66_63)), inference(subsumption_resolution, [], [f1813, f284])).
fof(f1813, plain, ((e0 = e2) | (~ spl66_61 | ~ spl66_63)), inference(forward_demodulation, [], [f855, f847])).
fof(f1810, plain, (~ spl66_18 | ~ spl66_3 | spl66_139), inference(avatar_split_clause, [], [f1809, f1467, f598, f662])).
fof(f1809, plain, (~ (e1 = op(e2, e3)) | (~ spl66_3 | spl66_139)), inference(forward_demodulation, [], [f1469, f600])).
fof(f1804, plain, (~ spl66_50 | ~ spl66_54), inference(avatar_split_clause, [], [f1803, f815, f798])).
fof(f1803, plain, (~ (e1 = op(e0, e3)) | ~ spl66_54), inference(forward_demodulation, [], [f264, f817])).
fof(f1798, plain, (~ spl66_42 | ~ spl66_10), inference(avatar_split_clause, [], [f1797, f628, f764])).
fof(f1797, plain, (~ (e1 = op(e1, e1)) | ~ spl66_10), inference(forward_demodulation, [], [f245, f630])).
fof(f630, plain, ((e1 = op(e3, e1)) | ~ spl66_10), inference(avatar_component_clause, [], [f628])).
fof(f1782, plain, (~ spl66_15 | ~ spl66_3), inference(avatar_split_clause, [], [f1781, f598, f649])).
fof(f1781, plain, (~ (e2 = op(e3, e0)) | ~ spl66_3), inference(forward_demodulation, [], [f279, f600])).
fof(f1778, plain, (~ spl66_13 | ~ spl66_61), inference(avatar_split_clause, [], [f1777, f845, f641])).
fof(f1777, plain, (~ (e0 = op(e3, e0)) | ~ spl66_61), inference(forward_demodulation, [], [f237, f847])).
fof(f1760, plain, (~ spl66_27 | ~ spl66_28), inference(avatar_contradiction_clause, [], [f1759])).
fof(f1759, plain, ($false | (~ spl66_27 | ~ spl66_28)), inference(subsumption_resolution, [], [f1758, f288])).
fof(f1758, plain, ((e2 = e3) | (~ spl66_27 | ~ spl66_28)), inference(forward_demodulation, [], [f706, f702])).
fof(f1720, plain, (~ spl66_59 | ~ spl66_27), inference(avatar_split_clause, [], [f1719, f700, f836])).
fof(f1719, plain, (~ (e2 = op(e0, e1)) | ~ spl66_27), inference(forward_demodulation, [], [f242, f702])).
fof(f1693, plain, (~ spl66_29 | ~ spl66_17), inference(avatar_split_clause, [], [f1692, f658, f709])).
fof(f1692, plain, (~ (e0 = op(e2, e0)) | ~ spl66_17), inference(forward_demodulation, [], [f273, f660])).
fof(f660, plain, ((e0 = op(e2, e3)) | ~ spl66_17), inference(avatar_component_clause, [], [f658])).
fof(f1687, plain, (~ spl66_21 | ~ spl66_3 | spl66_137), inference(avatar_split_clause, [], [f1681, f1458, f598, f675])).
fof(f1681, plain, (~ (e0 = op(e2, e2)) | (~ spl66_3 | spl66_137)), inference(backward_demodulation, [], [f1460, f600])).
fof(f1677, plain, (~ spl66_6 | ~ spl66_10), inference(avatar_split_clause, [], [f1675, f628, f611])).
fof(f1675, plain, (~ (e1 = op(e3, e2)) | ~ spl66_10), inference(backward_demodulation, [], [f280, f630])).
fof(f1668, plain, (~ spl66_17 | ~ spl66_19), inference(avatar_contradiction_clause, [], [f1667])).
fof(f1667, plain, ($false | (~ spl66_17 | ~ spl66_19)), inference(subsumption_resolution, [], [f1666, f284])).
fof(f1666, plain, ((e0 = e2) | (~ spl66_17 | ~ spl66_19)), inference(backward_demodulation, [], [f668, f660])).
fof(f1665, plain, (~ spl66_3 | ~ spl66_19), inference(avatar_split_clause, [], [f1664, f666, f598])).
fof(f1664, plain, (~ (e2 = op(e3, e3)) | ~ spl66_19), inference(backward_demodulation, [], [f258, f668])).
fof(f1663, plain, (~ spl66_23 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f1662])).
fof(f1662, plain, ($false | (~ spl66_23 | ~ spl66_24)), inference(subsumption_resolution, [], [f1661, f288])).
fof(f1661, plain, ((e2 = e3) | (~ spl66_23 | ~ spl66_24)), inference(backward_demodulation, [], [f689, f685])).
fof(f1660, plain, (~ spl66_5 | ~ spl66_24 | spl66_144), inference(avatar_split_clause, [], [f1656, f1491, f687, f607])).
fof(f1491, plain, (spl66_144 <=> (e0 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_144])])).
fof(f1656, plain, (~ (e0 = op(e3, e2)) | (~ spl66_24 | spl66_144)), inference(backward_demodulation, [], [f1493, f689])).
fof(f1493, plain, (~ (e0 = op(op(e2, e2), e2)) | spl66_144), inference(avatar_component_clause, [], [f1491])).
fof(f1659, plain, (~ spl66_1 | ~ spl66_24 | spl66_134), inference(avatar_split_clause, [], [f1655, f1444, f687, f590])).
fof(f1655, plain, (~ (e0 = op(e3, e3)) | (~ spl66_24 | spl66_134)), inference(backward_demodulation, [], [f1446, f689])).
fof(f1652, plain, (~ spl66_26 | ~ spl66_27), inference(avatar_contradiction_clause, [], [f1651])).
fof(f1651, plain, ($false | (~ spl66_26 | ~ spl66_27)), inference(subsumption_resolution, [], [f1650, f286])).
fof(f1650, plain, ((e1 = e2) | (~ spl66_26 | ~ spl66_27)), inference(backward_demodulation, [], [f702, f698])).
fof(f1635, plain, (~ spl66_34 | ~ spl66_35), inference(avatar_contradiction_clause, [], [f1634])).
fof(f1634, plain, ($false | (~ spl66_34 | ~ spl66_35)), inference(subsumption_resolution, [], [f1633, f286])).
fof(f1633, plain, ((e1 = e2) | (~ spl66_34 | ~ spl66_35)), inference(backward_demodulation, [], [f736, f732])).
fof(f1628, plain, (~ spl66_33 | ~ spl66_37), inference(avatar_split_clause, [], [f1625, f743, f726])).
fof(f1625, plain, (~ (e0 = op(e1, e3)) | ~ spl66_37), inference(backward_demodulation, [], [f270, f745])).
fof(f1618, plain, (~ spl66_33 | ~ spl66_41), inference(avatar_split_clause, [], [f1610, f760, f726])).
fof(f1610, plain, (~ (e0 = op(e1, e3)) | ~ spl66_41), inference(backward_demodulation, [], [f269, f762])).
fof(f1616, plain, (~ spl66_9 | ~ spl66_41), inference(avatar_split_clause, [], [f1608, f760, f624])).
fof(f1608, plain, (~ (e0 = op(e3, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f245, f762])).
fof(f1596, plain, (~ spl66_4 | ~ spl66_52), inference(avatar_split_clause, [], [f1593, f806, f602])).
fof(f1593, plain, (~ (e3 = op(e3, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f255, f808])).
fof(f1578, plain, (~ spl66_56 | ~ spl66_60), inference(avatar_split_clause, [], [f1573, f840, f823])).
fof(f1573, plain, (~ (e3 = op(e0, e2)) | ~ spl66_60), inference(backward_demodulation, [], [f262, f842])).
fof(f1575, plain, (~ spl66_44 | ~ spl66_60), inference(avatar_split_clause, [], [f1570, f840, f772])).
fof(f1570, plain, (~ (e3 = op(e1, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f241, f842])).
fof(f1569, plain, (~ spl66_61 | ~ spl66_62), inference(avatar_contradiction_clause, [], [f1568])).
fof(f1568, plain, ($false | (~ spl66_61 | ~ spl66_62)), inference(subsumption_resolution, [], [f1567, f283])).
fof(f1567, plain, ((e0 = e1) | (~ spl66_61 | ~ spl66_62)), inference(backward_demodulation, [], [f851, f847])).
fof(f1566, plain, (~ spl66_44 | ~ spl66_62 | spl66_152), inference(avatar_split_clause, [], [f1558, f1532, f849, f772])).
fof(f1532, plain, (spl66_152 <=> (e3 = op(op(e0, e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_152])])).
fof(f1558, plain, (~ (e3 = op(e1, e1)) | (~ spl66_62 | spl66_152)), inference(backward_demodulation, [], [f1534, f851])).
fof(f1534, plain, (~ (e3 = op(op(e0, e0), op(e0, e0))) | spl66_152), inference(avatar_component_clause, [], [f1532])).
fof(f1565, plain, (~ spl66_48 | ~ spl66_62 | spl66_141), inference(avatar_split_clause, [], [f1557, f1476, f849, f789])).
fof(f1557, plain, (~ (e3 = op(e1, e0)) | (~ spl66_62 | spl66_141)), inference(backward_demodulation, [], [f1478, f851])).
fof(f1562, plain, (~ spl66_58 | ~ spl66_62), inference(avatar_split_clause, [], [f1552, f849, f832])).
fof(f1552, plain, (~ (e1 = op(e0, e1)) | ~ spl66_62), inference(backward_demodulation, [], [f259, f851])).
fof(f1559, plain, (~ spl66_46 | ~ spl66_62), inference(avatar_split_clause, [], [f1549, f849, f781])).
fof(f1549, plain, (~ (e1 = op(e1, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f235, f851])).
fof(f1541, plain, (~ spl66_153 | ~ spl66_41 | ~ spl66_133), inference(avatar_split_clause, [], [f585, f1439, f760, f1538])).
fof(f585, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(op(e1, e1), e1)) & (e0 = op(e1, e1)) & (e3 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax26)).
fof(f1535, plain, (~ spl66_152 | ~ spl66_62 | ~ spl66_142), inference(avatar_split_clause, [], [f583, f1481, f849, f1532])).
fof(f583, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e2 = op(op(e0, e0), e0)) & (op(e0, e0) = e1) & (e3 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax24)).
fof(f1530, plain, (~ spl66_151 | ~ spl66_2 | ~ spl66_146), inference(avatar_split_clause, [], [f582, f1501, f594, f1526])).
fof(f582, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(op(e3, e3), e3)) & (e1 = op(e3, e3)) & (e2 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax23)).
fof(f1529, plain, (~ spl66_151 | ~ spl66_1 | ~ spl66_139), inference(avatar_split_clause, [], [f581, f1467, f590, f1526])).
fof(f581, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(op(e3, e3), e3)) & (e0 = op(e3, e3)) & (e2 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax22)).
fof(f1524, plain, (~ spl66_149 | ~ spl66_44 | ~ spl66_150), inference(avatar_split_clause, [], [f580, f1521, f772, f1516])).
fof(f580, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(op(e1, e1), e1)) & (e3 = op(e1, e1)) & (e2 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax21)).
fof(f1519, plain, (~ spl66_149 | ~ spl66_41 | ~ spl66_132), inference(avatar_split_clause, [], [f579, f1434, f760, f1516])).
fof(f579, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e3 = op(op(e1, e1), e1)) & (e0 = op(e1, e1)) & (e2 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax20)).
fof(f1514, plain, (~ spl66_147 | ~ spl66_64 | ~ spl66_148), inference(avatar_split_clause, [], [f578, f1511, f857, f1506])).
fof(f578, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(op(e0, e0), e0)) & (op(e0, e0) = e3) & (e2 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax19)).
fof(f1509, plain, (~ spl66_147 | ~ spl66_62 | ~ spl66_141), inference(avatar_split_clause, [], [f577, f1476, f849, f1506])).
fof(f577, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(op(e0, e0), e0)) & (op(e0, e0) = e1) & (e2 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax18)).
fof(f1499, plain, (~ spl66_145 | ~ spl66_1 | ~ spl66_138), inference(avatar_split_clause, [], [f575, f1462, f590, f1496])).
fof(f575, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e2 = op(op(e3, e3), e3)) & (e0 = op(e3, e3)) & (e1 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax16)).
fof(f1494, plain, (~ spl66_143 | ~ spl66_24 | ~ spl66_144), inference(avatar_split_clause, [], [f574, f1491, f687, f1486])).
fof(f574, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(op(e2, e2), e2)) & (e3 = op(e2, e2)) & (e1 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax15)).
fof(f1489, plain, (~ spl66_143 | ~ spl66_21 | ~ spl66_135), inference(avatar_split_clause, [], [f573, f1448, f675, f1486])).
fof(f573, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(op(e2, e2), e2)) & (e0 = op(e2, e2)) & (e1 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax14)).
fof(f1484, plain, (~ spl66_140 | ~ spl66_64 | ~ spl66_142), inference(avatar_split_clause, [], [f572, f1481, f857, f1472])).
fof(f572, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e2 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e2 = op(op(e0, e0), e0)) & (op(e0, e0) = e3) & (e1 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax13)).
fof(f1479, plain, (~ spl66_140 | ~ spl66_63 | ~ spl66_141), inference(avatar_split_clause, [], [f571, f1476, f853, f1472])).
fof(f571, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), op(e0, e0)))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e3 = op(op(e0, e0), e0)) | ~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), op(e0, e0)))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(op(e0, e0), e0)) & (op(e0, e0) = e2) & (e1 = op(op(e0, e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax12)).
fof(f1470, plain, (~ spl66_137 | ~ spl66_3 | ~ spl66_139), inference(avatar_split_clause, [], [f570, f1467, f598, f1458])).
fof(f570, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e1 = op(op(e3, e3), e3)) & (e2 = op(e3, e3)) & (e0 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax11)).
fof(f1465, plain, (~ spl66_137 | ~ spl66_2 | ~ spl66_138), inference(avatar_split_clause, [], [f569, f1462, f594, f1458])).
fof(f569, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), op(e3, e3)))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(op(e3, e3), e3)) & (e1 = op(e3, e3)) & (e0 = op(op(e3, e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax10)).
fof(f1456, plain, (~ spl66_134 | ~ spl66_24 | ~ spl66_136), inference(avatar_split_clause, [], [f568, f1453, f687, f1444])).
fof(f568, plain, (~ (e1 = op(op(e2, e2), e2)) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e1 = op(op(e2, e2), e2)) | ~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e1 = op(op(e2, e2), e2)) & (e3 = op(e2, e2)) & (e0 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax9)).
fof(f1451, plain, (~ spl66_134 | ~ spl66_22 | ~ spl66_135), inference(avatar_split_clause, [], [f567, f1448, f679, f1444])).
fof(f567, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), op(e2, e2)))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(op(e2, e2), e2)) & (e1 = op(e2, e2)) & (e0 = op(op(e2, e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax8)).
fof(f1442, plain, (~ spl66_131 | ~ spl66_44 | ~ spl66_133), inference(avatar_split_clause, [], [f566, f1439, f772, f1430])).
fof(f566, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e2 = op(op(e1, e1), e1)) & (e3 = op(e1, e1)) & (e0 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax7)).
fof(f1437, plain, (~ spl66_131 | ~ spl66_43 | ~ spl66_132), inference(avatar_split_clause, [], [f565, f1434, f768, f1430])).
fof(f565, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), op(e1, e1)))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e3 = op(op(e1, e1), e1)) & (e2 = op(e1, e1)) & (e0 = op(op(e1, e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax6)).
fof(f1428, plain, (spl66_130 | spl66_129 | spl66_128 | spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | spl66_79 | spl66_78 | spl66_77 | spl66_76 | spl66_75 | spl66_74 | spl66_73 | spl66_72 | spl66_71 | spl66_70 | spl66_69 | spl66_68 | spl66_4), inference(avatar_split_clause, [], [f556, f602, f918, f926, f934, f942, f950, f958, f966, f974, f982, f990, f998, f1006, f1014, f1022, f1030, f1038, f1046, f1054, f1062, f1070, f1078, f1086, f1094, f1102, f1110, f1118, f1126, f1134, f1142, f1150, f1158, f1166, f1174, f1182, f1190, f1198, f1206, f1214, f1222, f1230, f1238, f1246, f1254, f1262, f1270, f1278, f1286, f1294, f1302, f1310, f1318, f1326, f1334, f1342, f1350, f1358, f1366, f1374, f1382, f1390, f1398, f1406, f1414])).
fof(f1414, plain, (spl66_130 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl66_130])])).
fof(f1406, plain, (spl66_129 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl66_129])])).
fof(f1398, plain, (spl66_128 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl66_128])])).
fof(f1390, plain, (spl66_127 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl66_127])])).
fof(f1382, plain, (spl66_126 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl66_126])])).
fof(f1374, plain, (spl66_125 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl66_125])])).
fof(f1366, plain, (spl66_124 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl66_124])])).
fof(f1358, plain, (spl66_123 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl66_123])])).
fof(f1350, plain, (spl66_122 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl66_122])])).
fof(f1342, plain, (spl66_121 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl66_121])])).
fof(f1334, plain, (spl66_120 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl66_120])])).
fof(f1326, plain, (spl66_119 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl66_119])])).
fof(f1318, plain, (spl66_118 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl66_118])])).
fof(f1310, plain, (spl66_117 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl66_117])])).
fof(f1302, plain, (spl66_116 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl66_116])])).
fof(f1294, plain, (spl66_115 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl66_115])])).
fof(f1286, plain, (spl66_114 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl66_114])])).
fof(f1278, plain, (spl66_113 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl66_113])])).
fof(f1270, plain, (spl66_112 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl66_112])])).
fof(f1262, plain, (spl66_111 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl66_111])])).
fof(f1254, plain, (spl66_110 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl66_110])])).
fof(f1246, plain, (spl66_109 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl66_109])])).
fof(f1238, plain, (spl66_108 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl66_108])])).
fof(f1230, plain, (spl66_107 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl66_107])])).
fof(f1222, plain, (spl66_106 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl66_106])])).
fof(f1214, plain, (spl66_105 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl66_105])])).
fof(f1206, plain, (spl66_104 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl66_104])])).
fof(f1198, plain, (spl66_103 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl66_103])])).
fof(f1190, plain, (spl66_102 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl66_102])])).
fof(f1182, plain, (spl66_101 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl66_101])])).
fof(f1174, plain, (spl66_100 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl66_100])])).
fof(f1166, plain, (spl66_99 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl66_99])])).
fof(f1158, plain, (spl66_98 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl66_98])])).
fof(f1150, plain, (spl66_97 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl66_97])])).
fof(f1142, plain, (spl66_96 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl66_96])])).
fof(f1134, plain, (spl66_95 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl66_95])])).
fof(f1126, plain, (spl66_94 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl66_94])])).
fof(f1118, plain, (spl66_93 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl66_93])])).
fof(f1110, plain, (spl66_92 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl66_92])])).
fof(f1102, plain, (spl66_91 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl66_91])])).
fof(f1094, plain, (spl66_90 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl66_90])])).
fof(f1086, plain, (spl66_89 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl66_89])])).
fof(f1078, plain, (spl66_88 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl66_88])])).
fof(f1070, plain, (spl66_87 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl66_87])])).
fof(f1062, plain, (spl66_86 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl66_86])])).
fof(f1054, plain, (spl66_85 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl66_85])])).
fof(f1046, plain, (spl66_84 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl66_84])])).
fof(f1038, plain, (spl66_83 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl66_83])])).
fof(f1030, plain, (spl66_82 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl66_82])])).
fof(f1022, plain, (spl66_81 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl66_81])])).
fof(f1014, plain, (spl66_80 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl66_80])])).
fof(f1006, plain, (spl66_79 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl66_79])])).
fof(f998, plain, (spl66_78 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl66_78])])).
fof(f990, plain, (spl66_77 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl66_77])])).
fof(f982, plain, (spl66_76 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl66_76])])).
fof(f974, plain, (spl66_75 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl66_75])])).
fof(f966, plain, (spl66_74 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl66_74])])).
fof(f958, plain, (spl66_73 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl66_73])])).
fof(f950, plain, (spl66_72 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl66_72])])).
fof(f942, plain, (spl66_71 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl66_71])])).
fof(f934, plain, (spl66_70 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl66_70])])).
fof(f926, plain, (spl66_69 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl66_69])])).
fof(f918, plain, (spl66_68 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl66_68])])).
fof(f556, plain, ((e3 = op(e3, e3)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f120, plain, (((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e3 = op(e3, e3))) | sP65 | sP64 | sP63) & (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, (((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> ((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, (((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> ((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, (((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> ((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, (((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> ((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, (((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> ((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, (((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> ((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, (((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> ((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, (((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> ((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, (((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> ((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, (((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> ((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, (((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> ((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, (((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> ((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, (((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> ((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, (((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> ((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, (((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> ((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, (((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> ((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, (((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> ((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, (((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> ((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, (((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> ((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, (((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> ((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, (((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> ((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, (((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> ((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f84, plain, (((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP30), inference(usedef, [], [e84])).
fof(e84, plain, (sP30 <=> ((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f85, plain, (((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(usedef, [], [e85])).
fof(e85, plain, (sP31 <=> ((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f86, plain, (((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP32), inference(usedef, [], [e86])).
fof(e86, plain, (sP32 <=> ((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f87, plain, (((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP33), inference(usedef, [], [e87])).
fof(e87, plain, (sP33 <=> ((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f88, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP34), inference(usedef, [], [e88])).
fof(e88, plain, (sP34 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f89, plain, (((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(usedef, [], [e89])).
fof(e89, plain, (sP35 <=> ((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f90, plain, (((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP36), inference(usedef, [], [e90])).
fof(e90, plain, (sP36 <=> ((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f91, plain, (((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP37), inference(usedef, [], [e91])).
fof(e91, plain, (sP37 <=> ((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f92, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP38), inference(usedef, [], [e92])).
fof(e92, plain, (sP38 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f93, plain, (((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(usedef, [], [e93])).
fof(e93, plain, (sP39 <=> ((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f94, plain, (((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP40), inference(usedef, [], [e94])).
fof(e94, plain, (sP40 <=> ((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f95, plain, (((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP41), inference(usedef, [], [e95])).
fof(e95, plain, (sP41 <=> ((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f96, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP42), inference(usedef, [], [e96])).
fof(e96, plain, (sP42 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f97, plain, (((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(usedef, [], [e97])).
fof(e97, plain, (sP43 <=> ((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f98, plain, (((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP44), inference(usedef, [], [e98])).
fof(e98, plain, (sP44 <=> ((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f99, plain, (((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP45), inference(usedef, [], [e99])).
fof(e99, plain, (sP45 <=> ((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f100, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP46), inference(usedef, [], [e100])).
fof(e100, plain, (sP46 <=> ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f101, plain, (((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(usedef, [], [e101])).
fof(e101, plain, (sP47 <=> ((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f102, plain, (((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP48), inference(usedef, [], [e102])).
fof(e102, plain, (sP48 <=> ((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f103, plain, (((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP49), inference(usedef, [], [e103])).
fof(e103, plain, (sP49 <=> ((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f104, plain, (((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP50), inference(usedef, [], [e104])).
fof(e104, plain, (sP50 <=> ((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f105, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(usedef, [], [e105])).
fof(e105, plain, (sP51 <=> ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f106, plain, (((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP52), inference(usedef, [], [e106])).
fof(e106, plain, (sP52 <=> ((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f107, plain, (((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP53), inference(usedef, [], [e107])).
fof(e107, plain, (sP53 <=> ((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f108, plain, (((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP54), inference(usedef, [], [e108])).
fof(e108, plain, (sP54 <=> ((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f109, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(usedef, [], [e109])).
fof(e109, plain, (sP55 <=> ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f110, plain, (((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP56), inference(usedef, [], [e110])).
fof(e110, plain, (sP56 <=> ((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f111, plain, (((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP57), inference(usedef, [], [e111])).
fof(e111, plain, (sP57 <=> ((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f112, plain, (((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP58), inference(usedef, [], [e112])).
fof(e112, plain, (sP58 <=> ((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f113, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(usedef, [], [e113])).
fof(e113, plain, (sP59 <=> ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f114, plain, (((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP60), inference(usedef, [], [e114])).
fof(e114, plain, (sP60 <=> ((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f115, plain, (((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP61), inference(usedef, [], [e115])).
fof(e115, plain, (sP61 <=> ((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f116, plain, (((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP62), inference(usedef, [], [e116])).
fof(e116, plain, (sP62 <=> ((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f117, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0))) | ~ sP63), inference(usedef, [], [e117])).
fof(e117, plain, (sP63 <=> (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f118, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e1 = op(e1, e1))) | ~ sP64), inference(usedef, [], [e118])).
fof(e118, plain, (sP64 <=> (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f119, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e2 = op(e2, e2))) | ~ sP65), inference(usedef, [], [e119])).
fof(e119, plain, (sP65 <=> (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f5, plain, (((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3))) & ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3))) & ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3))) & (e3 = op(e3, e3))) | (((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e2 = op(e2, e2))) | (((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e1 = op(e1, e1))) | (((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)))) & (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax5)).
fof(f1427, plain, (spl66_130 | spl66_129 | spl66_128 | spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | spl66_79 | spl66_78 | spl66_77 | spl66_76 | spl66_75 | spl66_74 | spl66_73 | spl66_72 | spl66_71 | spl66_70 | spl66_69 | spl66_68 | ~ spl66_4), inference(avatar_split_clause, [], [f557, f602, f918, f926, f934, f942, f950, f958, f966, f974, f982, f990, f998, f1006, f1014, f1022, f1030, f1038, f1046, f1054, f1062, f1070, f1078, f1086, f1094, f1102, f1110, f1118, f1126, f1134, f1142, f1150, f1158, f1166, f1174, f1182, f1190, f1198, f1206, f1214, f1222, f1230, f1238, f1246, f1254, f1262, f1270, f1278, f1286, f1294, f1302, f1310, f1318, f1326, f1334, f1342, f1350, f1358, f1366, f1374, f1382, f1390, f1398, f1406, f1414])).
fof(f557, plain, (~ (e3 = op(e3, e3)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1424, plain, (spl66_67 | spl66_66 | spl66_65 | spl66_4), inference(avatar_split_clause, [], [f560, f602, f894, f902, f910])).
fof(f910, plain, (spl66_67 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl66_67])])).
fof(f902, plain, (spl66_66 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl66_66])])).
fof(f894, plain, (spl66_65 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl66_65])])).
fof(f560, plain, ((e3 = op(e3, e3)) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1423, plain, (spl66_67 | spl66_66 | spl66_65 | ~ spl66_49 | spl66_64), inference(avatar_split_clause, [], [f561, f857, f794, f894, f902, f910])).
fof(f561, plain, ((op(e0, e0) = e3) | ~ (e0 = op(e0, e3)) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1422, plain, (spl66_67 | spl66_66 | spl66_65 | ~ spl66_34 | spl66_44), inference(avatar_split_clause, [], [f562, f772, f730, f894, f902, f910])).
fof(f562, plain, ((e3 = op(e1, e1)) | ~ (e1 = op(e1, e3)) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1421, plain, (spl66_67 | spl66_66 | spl66_65 | ~ spl66_19 | spl66_24), inference(avatar_split_clause, [], [f563, f687, f666, f894, f902, f910])).
fof(f563, plain, ((e3 = op(e2, e2)) | ~ (e2 = op(e2, e3)) | sP65 | sP64 | sP63), inference(cnf_transformation, [], [f120])).
fof(f1420, plain, (~ spl66_130 | spl66_61), inference(avatar_split_clause, [], [f552, f845, f1414])).
fof(f552, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f186, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1419, plain, (~ spl66_130 | ~ spl66_61), inference(avatar_split_clause, [], [f553, f845, f1414])).
fof(f553, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1412, plain, (~ spl66_129 | spl66_61), inference(avatar_split_clause, [], [f548, f845, f1406])).
fof(f548, plain, ((e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f185, plain, (((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1411, plain, (~ spl66_129 | ~ spl66_61), inference(avatar_split_clause, [], [f549, f845, f1406])).
fof(f549, plain, (~ (e0 = op(e0, e0)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1409, plain, (~ spl66_129 | spl66_46), inference(avatar_split_clause, [], [f551, f781, f1406])).
fof(f551, plain, ((e1 = op(e1, e0)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1404, plain, (~ spl66_128 | spl66_61), inference(avatar_split_clause, [], [f544, f845, f1398])).
fof(f544, plain, ((e0 = op(e0, e0)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f184, plain, (((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1403, plain, (~ spl66_128 | ~ spl66_61), inference(avatar_split_clause, [], [f545, f845, f1398])).
fof(f545, plain, (~ (e0 = op(e0, e0)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1401, plain, (~ spl66_128 | spl66_31), inference(avatar_split_clause, [], [f547, f717, f1398])).
fof(f547, plain, ((e2 = op(e2, e0)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1396, plain, (~ spl66_127 | spl66_61), inference(avatar_split_clause, [], [f540, f845, f1390])).
fof(f540, plain, ((e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f183, plain, (((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1395, plain, (~ spl66_127 | ~ spl66_61), inference(avatar_split_clause, [], [f541, f845, f1390])).
fof(f541, plain, (~ (e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f1393, plain, (~ spl66_127 | spl66_16), inference(avatar_split_clause, [], [f543, f653, f1390])).
fof(f543, plain, ((e3 = op(e3, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f1388, plain, (~ spl66_126 | spl66_57), inference(avatar_split_clause, [], [f536, f828, f1382])).
fof(f536, plain, ((e0 = op(e0, e1)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f182, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1386, plain, (~ spl66_126 | spl66_61), inference(avatar_split_clause, [], [f538, f845, f1382])).
fof(f538, plain, ((e0 = op(e0, e0)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f1380, plain, (~ spl66_125 | spl66_57), inference(avatar_split_clause, [], [f532, f828, f1374])).
fof(f532, plain, ((e0 = op(e0, e1)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f181, plain, (((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1378, plain, (~ spl66_125 | spl66_41), inference(avatar_split_clause, [], [f534, f760, f1374])).
fof(f534, plain, ((e0 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f1372, plain, (~ spl66_124 | spl66_57), inference(avatar_split_clause, [], [f528, f828, f1366])).
fof(f528, plain, ((e0 = op(e0, e1)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f180, plain, (((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1371, plain, (~ spl66_124 | ~ spl66_62), inference(avatar_split_clause, [], [f529, f849, f1366])).
fof(f529, plain, (~ (op(e0, e0) = e1) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1370, plain, (~ spl66_124 | spl66_21), inference(avatar_split_clause, [], [f530, f675, f1366])).
fof(f530, plain, ((e0 = op(e2, e2)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1369, plain, (~ spl66_124 | spl66_31), inference(avatar_split_clause, [], [f531, f717, f1366])).
fof(f531, plain, ((e2 = op(e2, e0)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1364, plain, (~ spl66_123 | spl66_57), inference(avatar_split_clause, [], [f524, f828, f1358])).
fof(f524, plain, ((e0 = op(e0, e1)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f179, plain, (((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1363, plain, (~ spl66_123 | ~ spl66_62), inference(avatar_split_clause, [], [f525, f849, f1358])).
fof(f525, plain, (~ (op(e0, e0) = e1) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1362, plain, (~ spl66_123 | spl66_1), inference(avatar_split_clause, [], [f526, f590, f1358])).
fof(f526, plain, ((e0 = op(e3, e3)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1361, plain, (~ spl66_123 | spl66_16), inference(avatar_split_clause, [], [f527, f653, f1358])).
fof(f527, plain, ((e3 = op(e3, e0)) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1356, plain, (~ spl66_122 | spl66_53), inference(avatar_split_clause, [], [f520, f811, f1350])).
fof(f520, plain, ((e0 = op(e0, e2)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f178, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1354, plain, (~ spl66_122 | spl66_61), inference(avatar_split_clause, [], [f522, f845, f1350])).
fof(f522, plain, ((e0 = op(e0, e0)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f1348, plain, (~ spl66_121 | spl66_53), inference(avatar_split_clause, [], [f516, f811, f1342])).
fof(f516, plain, ((e0 = op(e0, e2)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f177, plain, (((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1347, plain, (~ spl66_121 | ~ spl66_63), inference(avatar_split_clause, [], [f517, f853, f1342])).
fof(f517, plain, (~ (op(e0, e0) = e2) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1346, plain, (~ spl66_121 | spl66_41), inference(avatar_split_clause, [], [f518, f760, f1342])).
fof(f518, plain, ((e0 = op(e1, e1)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1345, plain, (~ spl66_121 | spl66_46), inference(avatar_split_clause, [], [f519, f781, f1342])).
fof(f519, plain, ((e1 = op(e1, e0)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1340, plain, (~ spl66_120 | spl66_53), inference(avatar_split_clause, [], [f512, f811, f1334])).
fof(f512, plain, ((e0 = op(e0, e2)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f176, plain, (((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f1338, plain, (~ spl66_120 | spl66_21), inference(avatar_split_clause, [], [f514, f675, f1334])).
fof(f514, plain, ((e0 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1337, plain, (~ spl66_120 | spl66_31), inference(avatar_split_clause, [], [f515, f717, f1334])).
fof(f515, plain, ((e2 = op(e2, e0)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1332, plain, (~ spl66_119 | spl66_53), inference(avatar_split_clause, [], [f508, f811, f1326])).
fof(f508, plain, ((e0 = op(e0, e2)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f175, plain, (((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f1331, plain, (~ spl66_119 | ~ spl66_63), inference(avatar_split_clause, [], [f509, f853, f1326])).
fof(f509, plain, (~ (op(e0, e0) = e2) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1330, plain, (~ spl66_119 | spl66_1), inference(avatar_split_clause, [], [f510, f590, f1326])).
fof(f510, plain, ((e0 = op(e3, e3)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1329, plain, (~ spl66_119 | spl66_16), inference(avatar_split_clause, [], [f511, f653, f1326])).
fof(f511, plain, ((e3 = op(e3, e0)) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1324, plain, (~ spl66_118 | spl66_49), inference(avatar_split_clause, [], [f504, f794, f1318])).
fof(f504, plain, ((e0 = op(e0, e3)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f174, plain, (((e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f1322, plain, (~ spl66_118 | spl66_61), inference(avatar_split_clause, [], [f506, f845, f1318])).
fof(f506, plain, ((e0 = op(e0, e0)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1316, plain, (~ spl66_117 | spl66_49), inference(avatar_split_clause, [], [f500, f794, f1310])).
fof(f500, plain, ((e0 = op(e0, e3)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f173, plain, (((e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f1315, plain, (~ spl66_117 | ~ spl66_64), inference(avatar_split_clause, [], [f501, f857, f1310])).
fof(f501, plain, (~ (op(e0, e0) = e3) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1314, plain, (~ spl66_117 | spl66_41), inference(avatar_split_clause, [], [f502, f760, f1310])).
fof(f502, plain, ((e0 = op(e1, e1)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1313, plain, (~ spl66_117 | spl66_46), inference(avatar_split_clause, [], [f503, f781, f1310])).
fof(f503, plain, ((e1 = op(e1, e0)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1308, plain, (~ spl66_116 | spl66_49), inference(avatar_split_clause, [], [f496, f794, f1302])).
fof(f496, plain, ((e0 = op(e0, e3)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f172, plain, (((e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f1307, plain, (~ spl66_116 | ~ spl66_64), inference(avatar_split_clause, [], [f497, f857, f1302])).
fof(f497, plain, (~ (op(e0, e0) = e3) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1306, plain, (~ spl66_116 | spl66_21), inference(avatar_split_clause, [], [f498, f675, f1302])).
fof(f498, plain, ((e0 = op(e2, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1305, plain, (~ spl66_116 | spl66_31), inference(avatar_split_clause, [], [f499, f717, f1302])).
fof(f499, plain, ((e2 = op(e2, e0)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1300, plain, (~ spl66_115 | spl66_49), inference(avatar_split_clause, [], [f492, f794, f1294])).
fof(f492, plain, ((e0 = op(e0, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f171, plain, (((e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f1299, plain, (~ spl66_115 | ~ spl66_64), inference(avatar_split_clause, [], [f493, f857, f1294])).
fof(f493, plain, (~ (op(e0, e0) = e3) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1298, plain, (~ spl66_115 | spl66_1), inference(avatar_split_clause, [], [f494, f590, f1294])).
fof(f494, plain, ((e0 = op(e3, e3)) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1292, plain, (~ spl66_114 | spl66_46), inference(avatar_split_clause, [], [f488, f781, f1286])).
fof(f488, plain, ((e1 = op(e1, e0)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f170, plain, (((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f1291, plain, (~ spl66_114 | ~ spl66_41), inference(avatar_split_clause, [], [f489, f760, f1286])).
fof(f489, plain, (~ (e0 = op(e1, e1)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1290, plain, (~ spl66_114 | spl66_62), inference(avatar_split_clause, [], [f490, f849, f1286])).
fof(f490, plain, ((op(e0, e0) = e1) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1284, plain, (~ spl66_113 | spl66_46), inference(avatar_split_clause, [], [f484, f781, f1278])).
fof(f484, plain, ((e1 = op(e1, e0)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f169, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f1282, plain, (~ spl66_113 | spl66_42), inference(avatar_split_clause, [], [f486, f764, f1278])).
fof(f486, plain, ((e1 = op(e1, e1)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1276, plain, (~ spl66_112 | spl66_46), inference(avatar_split_clause, [], [f480, f781, f1270])).
fof(f480, plain, ((e1 = op(e1, e0)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f168, plain, (((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f1275, plain, (~ spl66_112 | ~ spl66_41), inference(avatar_split_clause, [], [f481, f760, f1270])).
fof(f481, plain, (~ (e0 = op(e1, e1)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1274, plain, (~ spl66_112 | spl66_22), inference(avatar_split_clause, [], [f482, f679, f1270])).
fof(f482, plain, ((e1 = op(e2, e2)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1273, plain, (~ spl66_112 | spl66_27), inference(avatar_split_clause, [], [f483, f700, f1270])).
fof(f483, plain, ((e2 = op(e2, e1)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1268, plain, (~ spl66_111 | spl66_46), inference(avatar_split_clause, [], [f476, f781, f1262])).
fof(f476, plain, ((e1 = op(e1, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f167, plain, (((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f1267, plain, (~ spl66_111 | ~ spl66_41), inference(avatar_split_clause, [], [f477, f760, f1262])).
fof(f477, plain, (~ (e0 = op(e1, e1)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1266, plain, (~ spl66_111 | spl66_2), inference(avatar_split_clause, [], [f478, f594, f1262])).
fof(f478, plain, ((e1 = op(e3, e3)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1265, plain, (~ spl66_111 | spl66_12), inference(avatar_split_clause, [], [f479, f636, f1262])).
fof(f479, plain, ((e3 = op(e3, e1)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1260, plain, (~ spl66_110 | spl66_42), inference(avatar_split_clause, [], [f472, f764, f1254])).
fof(f472, plain, ((e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f166, plain, (((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f1259, plain, (~ spl66_110 | ~ spl66_42), inference(avatar_split_clause, [], [f473, f764, f1254])).
fof(f473, plain, (~ (e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1258, plain, (~ spl66_110 | spl66_62), inference(avatar_split_clause, [], [f474, f849, f1254])).
fof(f474, plain, ((op(e0, e0) = e1) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1252, plain, (~ spl66_109 | spl66_42), inference(avatar_split_clause, [], [f468, f764, f1246])).
fof(f468, plain, ((e1 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f165, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f1251, plain, (~ spl66_109 | ~ spl66_42), inference(avatar_split_clause, [], [f469, f764, f1246])).
fof(f469, plain, (~ (e1 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1244, plain, (~ spl66_108 | spl66_42), inference(avatar_split_clause, [], [f464, f764, f1238])).
fof(f464, plain, ((e1 = op(e1, e1)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f164, plain, (((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f1243, plain, (~ spl66_108 | ~ spl66_42), inference(avatar_split_clause, [], [f465, f764, f1238])).
fof(f465, plain, (~ (e1 = op(e1, e1)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1236, plain, (~ spl66_107 | spl66_42), inference(avatar_split_clause, [], [f460, f764, f1230])).
fof(f460, plain, ((e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f163, plain, (((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f1235, plain, (~ spl66_107 | ~ spl66_42), inference(avatar_split_clause, [], [f461, f764, f1230])).
fof(f461, plain, (~ (e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f1228, plain, (~ spl66_106 | spl66_38), inference(avatar_split_clause, [], [f456, f747, f1222])).
fof(f456, plain, ((e1 = op(e1, e2)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f162, plain, (((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f1227, plain, (~ spl66_106 | ~ spl66_43), inference(avatar_split_clause, [], [f457, f768, f1222])).
fof(f457, plain, (~ (e2 = op(e1, e1)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1226, plain, (~ spl66_106 | spl66_62), inference(avatar_split_clause, [], [f458, f849, f1222])).
fof(f458, plain, ((op(e0, e0) = e1) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1225, plain, (~ spl66_106 | spl66_57), inference(avatar_split_clause, [], [f459, f828, f1222])).
fof(f459, plain, ((e0 = op(e0, e1)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1220, plain, (~ spl66_105 | spl66_38), inference(avatar_split_clause, [], [f452, f747, f1214])).
fof(f452, plain, ((e1 = op(e1, e2)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f161, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f1218, plain, (~ spl66_105 | spl66_42), inference(avatar_split_clause, [], [f454, f764, f1214])).
fof(f454, plain, ((e1 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f1212, plain, (~ spl66_104 | spl66_38), inference(avatar_split_clause, [], [f448, f747, f1206])).
fof(f448, plain, ((e1 = op(e1, e2)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f160, plain, (((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f1210, plain, (~ spl66_104 | spl66_22), inference(avatar_split_clause, [], [f450, f679, f1206])).
fof(f450, plain, ((e1 = op(e2, e2)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1209, plain, (~ spl66_104 | spl66_27), inference(avatar_split_clause, [], [f451, f700, f1206])).
fof(f451, plain, ((e2 = op(e2, e1)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1204, plain, (~ spl66_103 | spl66_38), inference(avatar_split_clause, [], [f444, f747, f1198])).
fof(f444, plain, ((e1 = op(e1, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f159, plain, (((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f1203, plain, (~ spl66_103 | ~ spl66_43), inference(avatar_split_clause, [], [f445, f768, f1198])).
fof(f445, plain, (~ (e2 = op(e1, e1)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1202, plain, (~ spl66_103 | spl66_2), inference(avatar_split_clause, [], [f446, f594, f1198])).
fof(f446, plain, ((e1 = op(e3, e3)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1201, plain, (~ spl66_103 | spl66_12), inference(avatar_split_clause, [], [f447, f636, f1198])).
fof(f447, plain, ((e3 = op(e3, e1)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1196, plain, (~ spl66_102 | spl66_34), inference(avatar_split_clause, [], [f440, f730, f1190])).
fof(f440, plain, ((e1 = op(e1, e3)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f158, plain, (((e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f1195, plain, (~ spl66_102 | ~ spl66_44), inference(avatar_split_clause, [], [f441, f772, f1190])).
fof(f441, plain, (~ (e3 = op(e1, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1194, plain, (~ spl66_102 | spl66_62), inference(avatar_split_clause, [], [f442, f849, f1190])).
fof(f442, plain, ((op(e0, e0) = e1) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1193, plain, (~ spl66_102 | spl66_57), inference(avatar_split_clause, [], [f443, f828, f1190])).
fof(f443, plain, ((e0 = op(e0, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1188, plain, (~ spl66_101 | spl66_34), inference(avatar_split_clause, [], [f436, f730, f1182])).
fof(f436, plain, ((e1 = op(e1, e3)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f157, plain, (((e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f1186, plain, (~ spl66_101 | spl66_42), inference(avatar_split_clause, [], [f438, f764, f1182])).
fof(f438, plain, ((e1 = op(e1, e1)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f1180, plain, (~ spl66_100 | spl66_34), inference(avatar_split_clause, [], [f432, f730, f1174])).
fof(f432, plain, ((e1 = op(e1, e3)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f156, plain, (((e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP30), inference(nnf_transformation, [], [f84])).
fof(f1179, plain, (~ spl66_100 | ~ spl66_44), inference(avatar_split_clause, [], [f433, f772, f1174])).
fof(f433, plain, (~ (e3 = op(e1, e1)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1178, plain, (~ spl66_100 | spl66_22), inference(avatar_split_clause, [], [f434, f679, f1174])).
fof(f434, plain, ((e1 = op(e2, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1177, plain, (~ spl66_100 | spl66_27), inference(avatar_split_clause, [], [f435, f700, f1174])).
fof(f435, plain, ((e2 = op(e2, e1)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1172, plain, (~ spl66_99 | spl66_34), inference(avatar_split_clause, [], [f428, f730, f1166])).
fof(f428, plain, ((e1 = op(e1, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f155, plain, (((e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(nnf_transformation, [], [f85])).
fof(f1170, plain, (~ spl66_99 | spl66_2), inference(avatar_split_clause, [], [f430, f594, f1166])).
fof(f430, plain, ((e1 = op(e3, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1164, plain, (~ spl66_98 | spl66_31), inference(avatar_split_clause, [], [f424, f717, f1158])).
fof(f424, plain, ((e2 = op(e2, e0)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f154, plain, (((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP32), inference(nnf_transformation, [], [f86])).
fof(f1162, plain, (~ spl66_98 | spl66_63), inference(avatar_split_clause, [], [f426, f853, f1158])).
fof(f426, plain, ((op(e0, e0) = e2) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1156, plain, (~ spl66_97 | spl66_31), inference(avatar_split_clause, [], [f420, f717, f1150])).
fof(f420, plain, ((e2 = op(e2, e0)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f153, plain, (((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP33), inference(nnf_transformation, [], [f87])).
fof(f1155, plain, (~ spl66_97 | ~ spl66_21), inference(avatar_split_clause, [], [f421, f675, f1150])).
fof(f421, plain, (~ (e0 = op(e2, e2)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1154, plain, (~ spl66_97 | spl66_43), inference(avatar_split_clause, [], [f422, f768, f1150])).
fof(f422, plain, ((e2 = op(e1, e1)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1153, plain, (~ spl66_97 | spl66_38), inference(avatar_split_clause, [], [f423, f747, f1150])).
fof(f423, plain, ((e1 = op(e1, e2)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1148, plain, (~ spl66_96 | spl66_31), inference(avatar_split_clause, [], [f416, f717, f1142])).
fof(f416, plain, ((e2 = op(e2, e0)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f152, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP34), inference(nnf_transformation, [], [f88])).
fof(f1146, plain, (~ spl66_96 | spl66_23), inference(avatar_split_clause, [], [f418, f683, f1142])).
fof(f418, plain, ((e2 = op(e2, e2)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f1140, plain, (~ spl66_95 | spl66_31), inference(avatar_split_clause, [], [f412, f717, f1134])).
fof(f412, plain, ((e2 = op(e2, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f151, plain, (((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(nnf_transformation, [], [f89])).
fof(f1139, plain, (~ spl66_95 | ~ spl66_21), inference(avatar_split_clause, [], [f413, f675, f1134])).
fof(f413, plain, (~ (e0 = op(e2, e2)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1138, plain, (~ spl66_95 | spl66_3), inference(avatar_split_clause, [], [f414, f598, f1134])).
fof(f414, plain, ((e2 = op(e3, e3)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1137, plain, (~ spl66_95 | spl66_8), inference(avatar_split_clause, [], [f415, f619, f1134])).
fof(f415, plain, ((e3 = op(e3, e2)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1132, plain, (~ spl66_94 | spl66_27), inference(avatar_split_clause, [], [f408, f700, f1126])).
fof(f408, plain, ((e2 = op(e2, e1)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f150, plain, (((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP36), inference(nnf_transformation, [], [f90])).
fof(f1131, plain, (~ spl66_94 | ~ spl66_22), inference(avatar_split_clause, [], [f409, f679, f1126])).
fof(f409, plain, (~ (e1 = op(e2, e2)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1130, plain, (~ spl66_94 | spl66_63), inference(avatar_split_clause, [], [f410, f853, f1126])).
fof(f410, plain, ((op(e0, e0) = e2) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1129, plain, (~ spl66_94 | spl66_53), inference(avatar_split_clause, [], [f411, f811, f1126])).
fof(f411, plain, ((e0 = op(e0, e2)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1124, plain, (~ spl66_93 | spl66_27), inference(avatar_split_clause, [], [f404, f700, f1118])).
fof(f404, plain, ((e2 = op(e2, e1)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f149, plain, (((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP37), inference(nnf_transformation, [], [f91])).
fof(f1122, plain, (~ spl66_93 | spl66_43), inference(avatar_split_clause, [], [f406, f768, f1118])).
fof(f406, plain, ((e2 = op(e1, e1)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1116, plain, (~ spl66_92 | spl66_27), inference(avatar_split_clause, [], [f400, f700, f1110])).
fof(f400, plain, ((e2 = op(e2, e1)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f148, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP38), inference(nnf_transformation, [], [f92])).
fof(f1115, plain, (~ spl66_92 | ~ spl66_22), inference(avatar_split_clause, [], [f401, f679, f1110])).
fof(f401, plain, (~ (e1 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1114, plain, (~ spl66_92 | spl66_23), inference(avatar_split_clause, [], [f402, f683, f1110])).
fof(f402, plain, ((e2 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1108, plain, (~ spl66_91 | spl66_27), inference(avatar_split_clause, [], [f396, f700, f1102])).
fof(f396, plain, ((e2 = op(e2, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f147, plain, (((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(nnf_transformation, [], [f93])).
fof(f1107, plain, (~ spl66_91 | ~ spl66_22), inference(avatar_split_clause, [], [f397, f679, f1102])).
fof(f397, plain, (~ (e1 = op(e2, e2)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1106, plain, (~ spl66_91 | spl66_3), inference(avatar_split_clause, [], [f398, f598, f1102])).
fof(f398, plain, ((e2 = op(e3, e3)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1105, plain, (~ spl66_91 | spl66_8), inference(avatar_split_clause, [], [f399, f619, f1102])).
fof(f399, plain, ((e3 = op(e3, e2)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1100, plain, (~ spl66_90 | spl66_23), inference(avatar_split_clause, [], [f392, f683, f1094])).
fof(f392, plain, ((e2 = op(e2, e2)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f146, plain, (((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP40), inference(nnf_transformation, [], [f94])).
fof(f1099, plain, (~ spl66_90 | ~ spl66_23), inference(avatar_split_clause, [], [f393, f683, f1094])).
fof(f393, plain, (~ (e2 = op(e2, e2)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f1098, plain, (~ spl66_90 | spl66_63), inference(avatar_split_clause, [], [f394, f853, f1094])).
fof(f394, plain, ((op(e0, e0) = e2) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f1092, plain, (~ spl66_89 | spl66_23), inference(avatar_split_clause, [], [f388, f683, f1086])).
fof(f388, plain, ((e2 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f145, plain, (((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP41), inference(nnf_transformation, [], [f95])).
fof(f1091, plain, (~ spl66_89 | ~ spl66_23), inference(avatar_split_clause, [], [f389, f683, f1086])).
fof(f389, plain, (~ (e2 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1084, plain, (~ spl66_88 | spl66_23), inference(avatar_split_clause, [], [f384, f683, f1078])).
fof(f384, plain, ((e2 = op(e2, e2)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f144, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP42), inference(nnf_transformation, [], [f96])).
fof(f1083, plain, (~ spl66_88 | ~ spl66_23), inference(avatar_split_clause, [], [f385, f683, f1078])).
fof(f385, plain, (~ (e2 = op(e2, e2)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1076, plain, (~ spl66_87 | spl66_23), inference(avatar_split_clause, [], [f380, f683, f1070])).
fof(f380, plain, ((e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f143, plain, (((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(nnf_transformation, [], [f97])).
fof(f1075, plain, (~ spl66_87 | ~ spl66_23), inference(avatar_split_clause, [], [f381, f683, f1070])).
fof(f381, plain, (~ (e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1073, plain, (~ spl66_87 | spl66_8), inference(avatar_split_clause, [], [f383, f619, f1070])).
fof(f383, plain, ((e3 = op(e3, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1068, plain, (~ spl66_86 | spl66_19), inference(avatar_split_clause, [], [f376, f666, f1062])).
fof(f376, plain, ((e2 = op(e2, e3)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f142, plain, (((e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP44), inference(nnf_transformation, [], [f98])).
fof(f1067, plain, (~ spl66_86 | ~ spl66_24), inference(avatar_split_clause, [], [f377, f687, f1062])).
fof(f377, plain, (~ (e3 = op(e2, e2)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1066, plain, (~ spl66_86 | spl66_63), inference(avatar_split_clause, [], [f378, f853, f1062])).
fof(f378, plain, ((op(e0, e0) = e2) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1065, plain, (~ spl66_86 | spl66_53), inference(avatar_split_clause, [], [f379, f811, f1062])).
fof(f379, plain, ((e0 = op(e0, e2)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1060, plain, (~ spl66_85 | spl66_19), inference(avatar_split_clause, [], [f372, f666, f1054])).
fof(f372, plain, ((e2 = op(e2, e3)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f141, plain, (((e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP45), inference(nnf_transformation, [], [f99])).
fof(f1059, plain, (~ spl66_85 | ~ spl66_24), inference(avatar_split_clause, [], [f373, f687, f1054])).
fof(f373, plain, (~ (e3 = op(e2, e2)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1058, plain, (~ spl66_85 | spl66_43), inference(avatar_split_clause, [], [f374, f768, f1054])).
fof(f374, plain, ((e2 = op(e1, e1)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1057, plain, (~ spl66_85 | spl66_38), inference(avatar_split_clause, [], [f375, f747, f1054])).
fof(f375, plain, ((e1 = op(e1, e2)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1052, plain, (~ spl66_84 | spl66_19), inference(avatar_split_clause, [], [f368, f666, f1046])).
fof(f368, plain, ((e2 = op(e2, e3)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f140, plain, (((e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP46), inference(nnf_transformation, [], [f100])).
fof(f1050, plain, (~ spl66_84 | spl66_23), inference(avatar_split_clause, [], [f370, f683, f1046])).
fof(f370, plain, ((e2 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f1044, plain, (~ spl66_83 | spl66_19), inference(avatar_split_clause, [], [f364, f666, f1038])).
fof(f364, plain, ((e2 = op(e2, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f139, plain, (((e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(nnf_transformation, [], [f101])).
fof(f1042, plain, (~ spl66_83 | spl66_3), inference(avatar_split_clause, [], [f366, f598, f1038])).
fof(f366, plain, ((e2 = op(e3, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1036, plain, (~ spl66_82 | spl66_16), inference(avatar_split_clause, [], [f360, f653, f1030])).
fof(f360, plain, ((e3 = op(e3, e0)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f138, plain, (((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP48), inference(nnf_transformation, [], [f102])).
fof(f1034, plain, (~ spl66_82 | spl66_64), inference(avatar_split_clause, [], [f362, f857, f1030])).
fof(f362, plain, ((op(e0, e0) = e3) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1028, plain, (~ spl66_81 | spl66_16), inference(avatar_split_clause, [], [f356, f653, f1022])).
fof(f356, plain, ((e3 = op(e3, e0)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f137, plain, (((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP49), inference(nnf_transformation, [], [f103])).
fof(f1027, plain, (~ spl66_81 | ~ spl66_1), inference(avatar_split_clause, [], [f357, f590, f1022])).
fof(f357, plain, (~ (e0 = op(e3, e3)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1026, plain, (~ spl66_81 | spl66_44), inference(avatar_split_clause, [], [f358, f772, f1022])).
fof(f358, plain, ((e3 = op(e1, e1)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1025, plain, (~ spl66_81 | spl66_34), inference(avatar_split_clause, [], [f359, f730, f1022])).
fof(f359, plain, ((e1 = op(e1, e3)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1020, plain, (~ spl66_80 | spl66_16), inference(avatar_split_clause, [], [f352, f653, f1014])).
fof(f352, plain, ((e3 = op(e3, e0)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f136, plain, (((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP50), inference(nnf_transformation, [], [f104])).
fof(f1019, plain, (~ spl66_80 | ~ spl66_1), inference(avatar_split_clause, [], [f353, f590, f1014])).
fof(f353, plain, (~ (e0 = op(e3, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1018, plain, (~ spl66_80 | spl66_24), inference(avatar_split_clause, [], [f354, f687, f1014])).
fof(f354, plain, ((e3 = op(e2, e2)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1017, plain, (~ spl66_80 | spl66_19), inference(avatar_split_clause, [], [f355, f666, f1014])).
fof(f355, plain, ((e2 = op(e2, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1012, plain, (~ spl66_79 | spl66_16), inference(avatar_split_clause, [], [f348, f653, f1006])).
fof(f348, plain, ((e3 = op(e3, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f135, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(nnf_transformation, [], [f105])).
fof(f1010, plain, (~ spl66_79 | spl66_4), inference(avatar_split_clause, [], [f350, f602, f1006])).
fof(f350, plain, ((e3 = op(e3, e3)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f1004, plain, (~ spl66_78 | spl66_12), inference(avatar_split_clause, [], [f344, f636, f998])).
fof(f344, plain, ((e3 = op(e3, e1)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f134, plain, (((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP52), inference(nnf_transformation, [], [f106])).
fof(f1003, plain, (~ spl66_78 | ~ spl66_2), inference(avatar_split_clause, [], [f345, f594, f998])).
fof(f345, plain, (~ (e1 = op(e3, e3)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1002, plain, (~ spl66_78 | spl66_64), inference(avatar_split_clause, [], [f346, f857, f998])).
fof(f346, plain, ((op(e0, e0) = e3) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1001, plain, (~ spl66_78 | spl66_49), inference(avatar_split_clause, [], [f347, f794, f998])).
fof(f347, plain, ((e0 = op(e0, e3)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f996, plain, (~ spl66_77 | spl66_12), inference(avatar_split_clause, [], [f340, f636, f990])).
fof(f340, plain, ((e3 = op(e3, e1)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f133, plain, (((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP53), inference(nnf_transformation, [], [f107])).
fof(f994, plain, (~ spl66_77 | spl66_44), inference(avatar_split_clause, [], [f342, f772, f990])).
fof(f342, plain, ((e3 = op(e1, e1)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f988, plain, (~ spl66_76 | spl66_12), inference(avatar_split_clause, [], [f336, f636, f982])).
fof(f336, plain, ((e3 = op(e3, e1)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f132, plain, (((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP54), inference(nnf_transformation, [], [f108])).
fof(f987, plain, (~ spl66_76 | ~ spl66_2), inference(avatar_split_clause, [], [f337, f594, f982])).
fof(f337, plain, (~ (e1 = op(e3, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f986, plain, (~ spl66_76 | spl66_24), inference(avatar_split_clause, [], [f338, f687, f982])).
fof(f338, plain, ((e3 = op(e2, e2)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f985, plain, (~ spl66_76 | spl66_19), inference(avatar_split_clause, [], [f339, f666, f982])).
fof(f339, plain, ((e2 = op(e2, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f980, plain, (~ spl66_75 | spl66_12), inference(avatar_split_clause, [], [f332, f636, f974])).
fof(f332, plain, ((e3 = op(e3, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f131, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(nnf_transformation, [], [f109])).
fof(f978, plain, (~ spl66_75 | spl66_4), inference(avatar_split_clause, [], [f334, f602, f974])).
fof(f334, plain, ((e3 = op(e3, e3)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f972, plain, (~ spl66_74 | spl66_8), inference(avatar_split_clause, [], [f328, f619, f966])).
fof(f328, plain, ((e3 = op(e3, e2)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP56), inference(nnf_transformation, [], [f110])).
fof(f971, plain, (~ spl66_74 | ~ spl66_3), inference(avatar_split_clause, [], [f329, f598, f966])).
fof(f329, plain, (~ (e2 = op(e3, e3)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f970, plain, (~ spl66_74 | spl66_64), inference(avatar_split_clause, [], [f330, f857, f966])).
fof(f330, plain, ((op(e0, e0) = e3) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f969, plain, (~ spl66_74 | spl66_49), inference(avatar_split_clause, [], [f331, f794, f966])).
fof(f331, plain, ((e0 = op(e0, e3)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f964, plain, (~ spl66_73 | spl66_8), inference(avatar_split_clause, [], [f324, f619, f958])).
fof(f324, plain, ((e3 = op(e3, e2)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f129, plain, (((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP57), inference(nnf_transformation, [], [f111])).
fof(f963, plain, (~ spl66_73 | ~ spl66_3), inference(avatar_split_clause, [], [f325, f598, f958])).
fof(f325, plain, (~ (e2 = op(e3, e3)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f962, plain, (~ spl66_73 | spl66_44), inference(avatar_split_clause, [], [f326, f772, f958])).
fof(f326, plain, ((e3 = op(e1, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f961, plain, (~ spl66_73 | spl66_34), inference(avatar_split_clause, [], [f327, f730, f958])).
fof(f327, plain, ((e1 = op(e1, e3)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f956, plain, (~ spl66_72 | spl66_8), inference(avatar_split_clause, [], [f320, f619, f950])).
fof(f320, plain, ((e3 = op(e3, e2)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f128, plain, (((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP58), inference(nnf_transformation, [], [f112])).
fof(f954, plain, (~ spl66_72 | spl66_24), inference(avatar_split_clause, [], [f322, f687, f950])).
fof(f322, plain, ((e3 = op(e2, e2)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f953, plain, (~ spl66_72 | spl66_19), inference(avatar_split_clause, [], [f323, f666, f950])).
fof(f323, plain, ((e2 = op(e2, e3)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f948, plain, (~ spl66_71 | spl66_8), inference(avatar_split_clause, [], [f316, f619, f942])).
fof(f316, plain, ((e3 = op(e3, e2)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f127, plain, (((e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(nnf_transformation, [], [f113])).
fof(f946, plain, (~ spl66_71 | spl66_4), inference(avatar_split_clause, [], [f318, f602, f942])).
fof(f318, plain, ((e3 = op(e3, e3)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f940, plain, (~ spl66_70 | spl66_4), inference(avatar_split_clause, [], [f312, f602, f934])).
fof(f312, plain, ((e3 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f126, plain, (((e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP60), inference(nnf_transformation, [], [f114])).
fof(f939, plain, (~ spl66_70 | ~ spl66_4), inference(avatar_split_clause, [], [f313, f602, f934])).
fof(f313, plain, (~ (e3 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f932, plain, (~ spl66_69 | spl66_4), inference(avatar_split_clause, [], [f308, f602, f926])).
fof(f308, plain, ((e3 = op(e3, e3)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f125, plain, (((e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP61), inference(nnf_transformation, [], [f115])).
fof(f931, plain, (~ spl66_69 | ~ spl66_4), inference(avatar_split_clause, [], [f309, f602, f926])).
fof(f309, plain, (~ (e3 = op(e3, e3)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f924, plain, (~ spl66_68 | spl66_4), inference(avatar_split_clause, [], [f304, f602, f918])).
fof(f304, plain, ((e3 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f124, plain, (((e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP62), inference(nnf_transformation, [], [f116])).
fof(f923, plain, (~ spl66_68 | ~ spl66_4), inference(avatar_split_clause, [], [f305, f602, f918])).
fof(f305, plain, (~ (e3 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f916, plain, (~ spl66_67 | spl66_61), inference(avatar_split_clause, [], [f299, f845, f910])).
fof(f299, plain, ((e0 = op(e0, e0)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0))) & ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0))) & ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0))) | ~ sP63), inference(nnf_transformation, [], [f117])).
fof(f915, plain, (~ spl66_67 | ~ spl66_46 | spl66_41), inference(avatar_split_clause, [], [f301, f760, f781, f910])).
fof(f301, plain, ((e0 = op(e1, e1)) | ~ (e1 = op(e1, e0)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f914, plain, (~ spl66_67 | ~ spl66_31 | spl66_21), inference(avatar_split_clause, [], [f302, f675, f717, f910])).
fof(f302, plain, ((e0 = op(e2, e2)) | ~ (e2 = op(e2, e0)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f913, plain, (~ spl66_67 | ~ spl66_16 | spl66_1), inference(avatar_split_clause, [], [f303, f590, f653, f910])).
fof(f303, plain, ((e0 = op(e3, e3)) | ~ (e3 = op(e3, e0)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f908, plain, (~ spl66_66 | spl66_42), inference(avatar_split_clause, [], [f294, f764, f902])).
fof(f294, plain, ((e1 = op(e1, e1)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1))) & ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1))) & (e1 = op(e1, e1))) | ~ sP64), inference(nnf_transformation, [], [f118])).
fof(f907, plain, (~ spl66_66 | ~ spl66_57 | spl66_62), inference(avatar_split_clause, [], [f295, f849, f828, f902])).
fof(f295, plain, ((op(e0, e0) = e1) | ~ (e0 = op(e0, e1)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f906, plain, (~ spl66_66 | ~ spl66_27 | spl66_22), inference(avatar_split_clause, [], [f297, f679, f700, f902])).
fof(f297, plain, ((e1 = op(e2, e2)) | ~ (e2 = op(e2, e1)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f905, plain, (~ spl66_66 | ~ spl66_12 | spl66_2), inference(avatar_split_clause, [], [f298, f594, f636, f902])).
fof(f298, plain, ((e1 = op(e3, e3)) | ~ (e3 = op(e3, e1)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f900, plain, (~ spl66_65 | spl66_23), inference(avatar_split_clause, [], [f289, f683, f894])).
fof(f289, plain, ((e2 = op(e2, e2)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2))) & ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2))) & (e2 = op(e2, e2))) | ~ sP65), inference(nnf_transformation, [], [f119])).
fof(f899, plain, (~ spl66_65 | ~ spl66_53 | spl66_63), inference(avatar_split_clause, [], [f290, f853, f811, f894])).
fof(f290, plain, ((op(e0, e0) = e2) | ~ (e0 = op(e0, e2)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f898, plain, (~ spl66_65 | ~ spl66_38 | spl66_43), inference(avatar_split_clause, [], [f291, f768, f747, f894])).
fof(f291, plain, ((e2 = op(e1, e1)) | ~ (e1 = op(e1, e2)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f897, plain, (~ spl66_65 | ~ spl66_8 | spl66_3), inference(avatar_split_clause, [], [f293, f598, f619, f894])).
fof(f293, plain, ((e2 = op(e3, e3)) | ~ (e3 = op(e3, e2)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f892, plain, (spl66_61 | spl66_57 | spl66_53 | spl66_49), inference(avatar_split_clause, [], [f203, f794, f811, f828, f845])).
fof(f203, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax2)).
fof(f890, plain, (spl66_62 | spl66_58 | spl66_54 | spl66_50), inference(avatar_split_clause, [], [f205, f798, f815, f832, f849])).
fof(f205, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f889, plain, (spl66_62 | spl66_46 | spl66_30 | spl66_14), inference(avatar_split_clause, [], [f206, f645, f713, f781, f849])).
fof(f206, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f888, plain, (spl66_63 | spl66_59 | spl66_55 | spl66_51), inference(avatar_split_clause, [], [f207, f802, f819, f836, f853])).
fof(f207, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f887, plain, (spl66_63 | spl66_47 | spl66_31 | spl66_15), inference(avatar_split_clause, [], [f208, f649, f717, f785, f853])).
fof(f208, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f886, plain, (spl66_64 | spl66_60 | spl66_56 | spl66_52), inference(avatar_split_clause, [], [f209, f806, f823, f840, f857])).
fof(f209, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f885, plain, (spl66_64 | spl66_48 | spl66_32 | spl66_16), inference(avatar_split_clause, [], [f210, f653, f721, f789, f857])).
fof(f210, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f884, plain, (spl66_45 | spl66_41 | spl66_37 | spl66_33), inference(avatar_split_clause, [], [f211, f726, f743, f760, f777])).
fof(f211, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f883, plain, (spl66_57 | spl66_41 | spl66_25 | spl66_9), inference(avatar_split_clause, [], [f212, f624, f692, f760, f828])).
fof(f212, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f880, plain, (spl66_47 | spl66_43 | spl66_39 | spl66_35), inference(avatar_split_clause, [], [f215, f734, f751, f768, f785])).
fof(f215, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f878, plain, (spl66_48 | spl66_44 | spl66_40 | spl66_36), inference(avatar_split_clause, [], [f217, f738, f755, f772, f789])).
fof(f217, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f877, plain, (spl66_60 | spl66_44 | spl66_28 | spl66_12), inference(avatar_split_clause, [], [f218, f636, f704, f772, f840])).
fof(f218, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f876, plain, (spl66_29 | spl66_25 | spl66_21 | spl66_17), inference(avatar_split_clause, [], [f219, f658, f675, f692, f709])).
fof(f219, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f875, plain, (spl66_53 | spl66_37 | spl66_21 | spl66_5), inference(avatar_split_clause, [], [f220, f607, f675, f743, f811])).
fof(f220, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f874, plain, (spl66_30 | spl66_26 | spl66_22 | spl66_18), inference(avatar_split_clause, [], [f221, f662, f679, f696, f713])).
fof(f221, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f873, plain, (spl66_54 | spl66_38 | spl66_22 | spl66_6), inference(avatar_split_clause, [], [f222, f611, f679, f747, f815])).
fof(f222, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f872, plain, (spl66_31 | spl66_27 | spl66_23 | spl66_19), inference(avatar_split_clause, [], [f223, f666, f683, f700, f717])).
fof(f223, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f870, plain, (spl66_32 | spl66_28 | spl66_24 | spl66_20), inference(avatar_split_clause, [], [f225, f670, f687, f704, f721])).
fof(f225, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f869, plain, (spl66_56 | spl66_40 | spl66_24 | spl66_8), inference(avatar_split_clause, [], [f226, f619, f687, f755, f823])).
fof(f226, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f867, plain, (spl66_49 | spl66_33 | spl66_17 | spl66_1), inference(avatar_split_clause, [], [f228, f590, f658, f726, f794])).
fof(f228, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f866, plain, (spl66_14 | spl66_10 | spl66_6 | spl66_2), inference(avatar_split_clause, [], [f229, f594, f611, f628, f645])).
fof(f229, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f865, plain, (spl66_50 | spl66_34 | spl66_18 | spl66_2), inference(avatar_split_clause, [], [f230, f594, f662, f730, f798])).
fof(f230, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f864, plain, (spl66_15 | spl66_11 | spl66_7 | spl66_3), inference(avatar_split_clause, [], [f231, f598, f615, f632, f649])).
fof(f231, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f863, plain, (spl66_51 | spl66_35 | spl66_19 | spl66_3), inference(avatar_split_clause, [], [f232, f598, f666, f734, f802])).
fof(f232, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f862, plain, (spl66_16 | spl66_12 | spl66_8 | spl66_4), inference(avatar_split_clause, [], [f233, f602, f619, f636, f653])).
fof(f233, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f861, plain, (spl66_52 | spl66_36 | spl66_20 | spl66_4), inference(avatar_split_clause, [], [f234, f602, f670, f738, f806])).
fof(f234, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f860, plain, (spl66_61 | spl66_62 | spl66_63 | spl66_64), inference(avatar_split_clause, [], [f187, f857, f853, f849, f845])).
fof(f187, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG162+1.p', ax1)).
fof(f843, plain, (spl66_57 | spl66_58 | spl66_59 | spl66_60), inference(avatar_split_clause, [], [f188, f840, f836, f832, f828])).
fof(f188, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f826, plain, (spl66_53 | spl66_54 | spl66_55 | spl66_56), inference(avatar_split_clause, [], [f189, f823, f819, f815, f811])).
fof(f189, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f809, plain, (spl66_49 | spl66_50 | spl66_51 | spl66_52), inference(avatar_split_clause, [], [f190, f806, f802, f798, f794])).
fof(f190, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f792, plain, (spl66_45 | spl66_46 | spl66_47 | spl66_48), inference(avatar_split_clause, [], [f191, f789, f785, f781, f777])).
fof(f191, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f775, plain, (spl66_41 | spl66_42 | spl66_43 | spl66_44), inference(avatar_split_clause, [], [f192, f772, f768, f764, f760])).
fof(f192, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f758, plain, (spl66_37 | spl66_38 | spl66_39 | spl66_40), inference(avatar_split_clause, [], [f193, f755, f751, f747, f743])).
fof(f193, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f741, plain, (spl66_33 | spl66_34 | spl66_35 | spl66_36), inference(avatar_split_clause, [], [f194, f738, f734, f730, f726])).
fof(f194, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f724, plain, (spl66_29 | spl66_30 | spl66_31 | spl66_32), inference(avatar_split_clause, [], [f195, f721, f717, f713, f709])).
fof(f195, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f707, plain, (spl66_25 | spl66_26 | spl66_27 | spl66_28), inference(avatar_split_clause, [], [f196, f704, f700, f696, f692])).
fof(f196, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f690, plain, (spl66_21 | spl66_22 | spl66_23 | spl66_24), inference(avatar_split_clause, [], [f197, f687, f683, f679, f675])).
fof(f197, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f673, plain, (spl66_17 | spl66_18 | spl66_19 | spl66_20), inference(avatar_split_clause, [], [f198, f670, f666, f662, f658])).
fof(f198, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f656, plain, (spl66_13 | spl66_14 | spl66_15 | spl66_16), inference(avatar_split_clause, [], [f199, f653, f649, f645, f641])).
fof(f199, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f639, plain, (spl66_9 | spl66_10 | spl66_11 | spl66_12), inference(avatar_split_clause, [], [f200, f636, f632, f628, f624])).
fof(f200, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f622, plain, (spl66_5 | spl66_6 | spl66_7 | spl66_8), inference(avatar_split_clause, [], [f201, f619, f615, f611, f607])).
fof(f201, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f605, plain, (spl66_1 | spl66_2 | spl66_3 | spl66_4), inference(avatar_split_clause, [], [f202, f602, f598, f594, f590])).
fof(f202, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).