fof(f5597, plain, $false, inference(avatar_sat_refutation, [], [f449, f466, f483, f500, f517, f534, f551, f568, f585, f602, f619, f636, f653, f670, f687, f704, f705, f706, f707, f708, f709, f710, f711, f712, f713, f714, f715, f716, f717, f718, f719, f720, f721, f722, f723, f724, f725, f726, f727, f728, f729, f730, f731, f732, f733, f734, f735, f736, f745, f750, f755, f760, f761, f762, f771, f776, f781, f786, f787, f788, f797, f802, f807, f812, f813, f814, f823, f828, f833, f838, f839, f840, f849, f850, f855, f856, f857, f858, f859, f860, f865, f866, f867, f868, f869, f870, f875, f876, f877, f878, f879, f880, f885, f886, f887, f888, f889, f890, f898, f899, f900, f905, f906, f907, f908, f909, f910, f915, f916, f917, f918, f919, f920, f925, f926, f927, f928, f929, f930, f935, f936, f937, f938, f939, f940, f945, f946, f949, f950, f955, f956, f957, f958, f959, f960, f965, f966, f967, f968, f969, f970, f975, f976, f977, f978, f979, f980, f985, f986, f987, f988, f989, f990, f999, f1000, f1005, f1006, f1007, f1008, f1009, f1010, f1015, f1016, f1017, f1018, f1019, f1020, f1025, f1026, f1027, f1028, f1029, f1030, f1035, f1036, f1037, f1038, f1039, f1040, f1049, f1050, f1055, f1056, f1057, f1058, f1059, f1060, f1065, f1066, f1067, f1068, f1069, f1070, f1075, f1076, f1077, f1078, f1079, f1080, f1085, f1086, f1087, f1088, f1089, f1090, f1099, f1100, f1101, f1102, f1103, f1104, f1105, f1106, f1107, f1108, f1109, f1110, f1111, f1112, f1121, f1126, f1135, f1136, f1141, f1142, f1151, f1156, f1165, f1166, f1171, f1172, f1181, f1186, f1195, f1196, f1201, f1202, f1211, f1216, f1225, f1226, f1231, f1232, f1263, f1276, f1300, f1313, f1320, f1328, f1330, f1333, f1346, f1355, f1388, f1394, f1407, f1408, f1413, f1424, f1427, f1450, f1467, f1479, f1490, f1503, f1508, f1511, f1524, f1534, f1547, f1559, f1581, f1599, f1604, f1613, f1624, f1634, f1640, f1648, f1656, f1691, f1717, f1722, f1753, f1772, f1774, f1793, f1804, f1827, f1847, f1850, f1857, f1860, f1875, f1877, f1900, f1931, f1940, f1949, f1966, f1969, f1987, f2009, f2018, f2022, f2028, f2037, f2039, f2093, f2130, f2133, f2137, f2155, f2180, f2187, f2190, f2199, f2213, f2225, f2232, f2240, f2244, f2252, f2277, f2282, f2301, f2314, f2327, f2346, f2428, f2437, f2459, f2510, f2521, f2556, f2564, f2580, f2583, f2594, f2630, f2645, f2678, f2686, f2708, f2736, f2748, f2802, f2847, f2874, f2894, f2910, f2960, f2967, f2969, f2973, f3006, f3017, f3018, f3032, f3044, f3053, f3082, f3093, f3103, f3122, f3148, f3176, f3196, f3213, f3220, f3239, f3241, f3245, f3264, f3296, f3301, f3315, f3325, f3359, f3411, f3423, f3425, f3431, f3443, f3465, f3474, f3478, f3491, f3528, f3540, f3559, f3599, f3603, f3607, f3615, f3635, f3649, f3650, f3683, f3711, f3713, f3738, f3742, f3764, f3788, f3803, f3830, f3868, f3880, f3904, f3923, f3944, f3989, f3995, f4003, f4004, f4039, f4042, f4049, f4059, f4076, f4077, f4088, f4095, f4096, f4104, f4174, f4180, f4196, f4209, f4216, f4230, f4238, f4270, f4282, f4313, f4319, f4370, f4392, f4407, f4440, f4441, f4450, f4455, f4459, f4472, f4473, f4492, f4495, f4504, f4528, f4529, f4540, f4541, f4546, f4547, f4555, f4568, f4579, f4580, f4588, f4593, f4602, f4603, f4604, f4612, f4633, f4637, f4650, f4652, f4654, f4658, f4659, f4681, f4687, f4692, f4705, f4719, f4723, f4737, f4747, f4754, f4756, f4761, f4763, f4778, f4788, f4789, f4801, f4808, f4809, f4815, f4827, f4839, f4840, f4844, f4853, f4858, f4876, f4877, f4882, f4895, f4916, f4917, f4924, f4925, f4931, f4937, f4938, f4939, f4947, f4948, f4959, f4960, f4967, f4984, f4989, f4992, f4996, f4999, f5001, f5028, f5029, f5030, f5039, f5051, f5052, f5073, f5100, f5101, f5102, f5116, f5122, f5132, f5133, f5134, f5150, f5151, f5159, f5169, f5172, f5218, f5220, f5226, f5242, f5243, f5251, f5257, f5258, f5264, f5274, f5281, f5288, f5309, f5324, f5325, f5326, f5337, f5338, f5339, f5347, f5353, f5361, f5362, f5363, f5378, f5380, f5386, f5391, f5392, f5398, f5404, f5405, f5411, f5412, f5418, f5426, f5430, f5456, f5463, f5486, f5487, f5499, f5500, f5506, f5512, f5520, f5521, f5533, f5540, f5550, f5551, f5556, f5565, f5569, f5575, f5576, f5582, f5594])).
fof(f5594, plain, (~ spl30_3 | ~ spl30_4 | spl30_131), inference(avatar_split_clause, [], [f5591, f1213, f446, f442])).
fof(f442, plain, (spl30_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_3])])).
fof(f446, plain, (spl30_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_4])])).
fof(f1213, plain, (spl30_131 <=> (e2 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_131])])).
fof(f5591, plain, (~ (e2 = op(e3, e3)) | (~ spl30_4 | spl30_131)), inference(backward_demodulation, [], [f1215, f448])).
fof(f448, plain, ((e3 = op(e3, e3)) | ~ spl30_4), inference(avatar_component_clause, [], [f446])).
fof(f1215, plain, (~ (e2 = op(op(e3, e3), e3)) | spl30_131), inference(avatar_component_clause, [], [f1213])).
fof(f5582, plain, (~ spl30_6 | ~ spl30_10), inference(avatar_split_clause, [], [f5577, f472, f455])).
fof(f455, plain, (spl30_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_6])])).
fof(f472, plain, (spl30_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_10])])).
fof(f5577, plain, (~ (e1 = op(e3, e2)) | ~ spl30_10), inference(backward_demodulation, [], [f208, f474])).
fof(f474, plain, ((e1 = op(e3, e1)) | ~ spl30_10), inference(avatar_component_clause, [], [f472])).
fof(f208, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax3)).
fof(f5576, plain, (spl30_5 | ~ spl30_15 | ~ spl30_84), inference(avatar_split_clause, [], [f5574, f835, f493, f451])).
fof(f451, plain, (spl30_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_5])])).
fof(f493, plain, (spl30_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_15])])).
fof(f835, plain, (spl30_84 <=> (e0 = op(e3, op(e3, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_84])])).
fof(f5574, plain, ((e0 = op(e3, e2)) | (~ spl30_15 | ~ spl30_84)), inference(backward_demodulation, [], [f836, f495])).
fof(f495, plain, ((e2 = op(e3, e0)) | ~ spl30_15), inference(avatar_component_clause, [], [f493])).
fof(f836, plain, ((e0 = op(e3, op(e3, e0))) | ~ spl30_84), inference(avatar_component_clause, [], [f835])).
fof(f5575, plain, (~ spl30_3 | ~ spl30_15), inference(avatar_split_clause, [], [f5572, f493, f442])).
fof(f5572, plain, (~ (e2 = op(e3, e3)) | ~ spl30_15), inference(backward_demodulation, [], [f207, f495])).
fof(f207, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5569, plain, (~ spl30_3 | ~ spl30_19), inference(avatar_split_clause, [], [f5567, f510, f442])).
fof(f510, plain, (spl30_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_19])])).
fof(f5567, plain, (~ (e2 = op(e3, e3)) | ~ spl30_19), inference(backward_demodulation, [], [f186, f512])).
fof(f512, plain, ((e2 = op(e2, e3)) | ~ spl30_19), inference(avatar_component_clause, [], [f510])).
fof(f186, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f5565, plain, (~ spl30_6 | ~ spl30_22), inference(avatar_split_clause, [], [f5557, f523, f455])).
fof(f523, plain, (spl30_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_22])])).
fof(f5557, plain, (~ (e1 = op(e3, e2)) | ~ spl30_22), inference(backward_demodulation, [], [f180, f525])).
fof(f525, plain, ((e1 = op(e2, e2)) | ~ spl30_22), inference(avatar_component_clause, [], [f523])).
fof(f180, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5556, plain, (~ spl30_12 | ~ spl30_28), inference(avatar_split_clause, [], [f5552, f548, f480])).
fof(f480, plain, (spl30_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_12])])).
fof(f548, plain, (spl30_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_28])])).
fof(f5552, plain, (~ (e3 = op(e3, e1)) | ~ spl30_28), inference(backward_demodulation, [], [f174, f550])).
fof(f550, plain, ((e3 = op(e2, e1)) | ~ spl30_28), inference(avatar_component_clause, [], [f548])).
fof(f174, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f5551, plain, (~ spl30_17 | ~ spl30_29), inference(avatar_split_clause, [], [f5545, f553, f502])).
fof(f502, plain, (spl30_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_17])])).
fof(f553, plain, (spl30_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_29])])).
fof(f5545, plain, (~ (e0 = op(e2, e3)) | ~ spl30_29), inference(backward_demodulation, [], [f201, f555])).
fof(f555, plain, ((e0 = op(e2, e0)) | ~ spl30_29), inference(avatar_component_clause, [], [f553])).
fof(f201, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5550, plain, (~ spl30_13 | ~ spl30_29), inference(avatar_split_clause, [], [f5542, f553, f485])).
fof(f485, plain, (spl30_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_13])])).
fof(f5542, plain, (~ (e0 = op(e3, e0)) | ~ spl30_29), inference(backward_demodulation, [], [f168, f555])).
fof(f168, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f5540, plain, (~ spl30_17 | ~ spl30_33), inference(avatar_split_clause, [], [f5534, f570, f502])).
fof(f570, plain, (spl30_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_33])])).
fof(f5534, plain, (~ (e0 = op(e2, e3)) | ~ spl30_33), inference(backward_demodulation, [], [f184, f572])).
fof(f572, plain, ((e0 = op(e1, e3)) | ~ spl30_33), inference(avatar_component_clause, [], [f570])).
fof(f184, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5533, plain, (~ spl30_5 | ~ spl30_43 | spl30_117 | ~ spl30_119), inference(avatar_split_clause, [], [f5529, f1153, f1144, f612, f451])).
fof(f612, plain, (spl30_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_43])])).
fof(f1144, plain, (spl30_117 <=> (e0 = op(op(op(e1, e1), e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_117])])).
fof(f1153, plain, (spl30_119 <=> (e3 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_119])])).
fof(f5529, plain, (~ (e0 = op(e3, e2)) | (~ spl30_43 | spl30_117 | ~ spl30_119)), inference(backward_demodulation, [], [f5469, f614])).
fof(f614, plain, ((e2 = op(e1, e1)) | ~ spl30_43), inference(avatar_component_clause, [], [f612])).
fof(f5469, plain, (~ (e0 = op(e3, op(e1, e1))) | (spl30_117 | ~ spl30_119)), inference(backward_demodulation, [], [f1146, f1154])).
fof(f1154, plain, ((e3 = op(op(e1, e1), e1)) | ~ spl30_119), inference(avatar_component_clause, [], [f1153])).
fof(f1146, plain, (~ (e0 = op(op(op(e1, e1), e1), op(e1, e1))) | spl30_117), inference(avatar_component_clause, [], [f1144])).
fof(f5521, plain, (~ spl30_41 | ~ spl30_46 | spl30_74), inference(avatar_split_clause, [], [f5519, f783, f625, f604])).
fof(f604, plain, (spl30_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_41])])).
fof(f625, plain, (spl30_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_46])])).
fof(f783, plain, (spl30_74 <=> (e0 = op(e1, op(e1, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_74])])).
fof(f5519, plain, (~ (e0 = op(e1, e1)) | (~ spl30_46 | spl30_74)), inference(backward_demodulation, [], [f785, f627])).
fof(f627, plain, ((e1 = op(e1, e0)) | ~ spl30_46), inference(avatar_component_clause, [], [f625])).
fof(f785, plain, (~ (e0 = op(e1, op(e1, e0))) | spl30_74), inference(avatar_component_clause, [], [f783])).
fof(f5520, plain, (~ spl30_34 | ~ spl30_46), inference(avatar_split_clause, [], [f5516, f625, f574])).
fof(f574, plain, (spl30_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_34])])).
fof(f5516, plain, (~ (e1 = op(e1, e3)) | ~ spl30_46), inference(backward_demodulation, [], [f195, f627])).
fof(f195, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f5512, plain, (~ spl30_34 | ~ spl30_50), inference(avatar_split_clause, [], [f5507, f642, f574])).
fof(f642, plain, (spl30_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_50])])).
fof(f5507, plain, (~ (e1 = op(e1, e3)) | ~ spl30_50), inference(backward_demodulation, [], [f181, f644])).
fof(f644, plain, ((e1 = op(e0, e3)) | ~ spl30_50), inference(avatar_component_clause, [], [f642])).
fof(f181, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f5506, plain, (~ spl30_23 | ~ spl30_55), inference(avatar_split_clause, [], [f5501, f663, f527])).
fof(f527, plain, (spl30_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_23])])).
fof(f663, plain, (spl30_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_55])])).
fof(f5501, plain, (~ (e2 = op(e2, e2)) | ~ spl30_55), inference(backward_demodulation, [], [f176, f665])).
fof(f665, plain, ((e2 = op(e0, e2)) | ~ spl30_55), inference(avatar_component_clause, [], [f663])).
fof(f176, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f5500, plain, (~ spl30_53 | ~ spl30_57), inference(avatar_split_clause, [], [f5493, f672, f655])).
fof(f655, plain, (spl30_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_53])])).
fof(f672, plain, (spl30_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_57])])).
fof(f5493, plain, (~ (e0 = op(e0, e2)) | ~ spl30_57), inference(backward_demodulation, [], [f190, f674])).
fof(f674, plain, ((e0 = op(e0, e1)) | ~ spl30_57), inference(avatar_component_clause, [], [f672])).
fof(f190, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f5499, plain, (~ spl30_41 | ~ spl30_57), inference(avatar_split_clause, [], [f5490, f672, f604])).
fof(f5490, plain, (~ (e0 = op(e1, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f169, f674])).
fof(f169, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f5487, plain, (~ spl30_52 | ~ spl30_64), inference(avatar_split_clause, [], [f5480, f701, f650])).
fof(f650, plain, (spl30_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_52])])).
fof(f701, plain, (spl30_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl30_64])])).
fof(f5480, plain, (~ (e3 = op(e0, e3)) | ~ spl30_64), inference(backward_demodulation, [], [f189, f703])).
fof(f703, plain, ((op(e0, e0) = e3) | ~ spl30_64), inference(avatar_component_clause, [], [f701])).
fof(f189, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f5486, plain, (~ spl30_32 | ~ spl30_64), inference(avatar_split_clause, [], [f5476, f701, f565])).
fof(f565, plain, (spl30_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_32])])).
fof(f5476, plain, (~ (e3 = op(e2, e0)) | ~ spl30_64), inference(backward_demodulation, [], [f164, f703])).
fof(f164, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f5463, plain, (~ spl30_24 | ~ spl30_40), inference(avatar_split_clause, [], [f5462, f599, f531])).
fof(f531, plain, (spl30_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_24])])).
fof(f599, plain, (spl30_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_40])])).
fof(f5462, plain, (~ (e3 = op(e2, e2)) | ~ spl30_40), inference(forward_demodulation, [], [f178, f601])).
fof(f601, plain, ((e3 = op(e1, e2)) | ~ spl30_40), inference(avatar_component_clause, [], [f599])).
fof(f178, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f5456, plain, (~ spl30_32 | ~ spl30_41 | ~ spl30_59 | spl30_122), inference(avatar_contradiction_clause, [], [f5455])).
fof(f5455, plain, ($false | (~ spl30_32 | ~ spl30_41 | ~ spl30_59 | spl30_122)), inference(subsumption_resolution, [], [f5454, f567])).
fof(f567, plain, ((e3 = op(e2, e0)) | ~ spl30_32), inference(avatar_component_clause, [], [f565])).
fof(f5454, plain, (~ (e3 = op(e2, e0)) | (~ spl30_41 | ~ spl30_59 | spl30_122)), inference(forward_demodulation, [], [f5453, f682])).
fof(f682, plain, ((e2 = op(e0, e1)) | ~ spl30_59), inference(avatar_component_clause, [], [f680])).
fof(f680, plain, (spl30_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_59])])).
fof(f5453, plain, (~ (e3 = op(op(e0, e1), e0)) | (~ spl30_41 | spl30_122)), inference(forward_demodulation, [], [f1170, f606])).
fof(f606, plain, ((e0 = op(e1, e1)) | ~ spl30_41), inference(avatar_component_clause, [], [f604])).
fof(f1170, plain, (~ (e3 = op(op(op(e1, e1), e1), op(e1, e1))) | spl30_122), inference(avatar_component_clause, [], [f1168])).
fof(f1168, plain, (spl30_122 <=> (e3 = op(op(op(e1, e1), e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_122])])).
fof(f5430, plain, (~ spl30_8 | ~ spl30_12), inference(avatar_split_clause, [], [f5427, f480, f463])).
fof(f463, plain, (spl30_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_8])])).
fof(f5427, plain, (~ (e3 = op(e3, e2)) | ~ spl30_12), inference(backward_demodulation, [], [f208, f482])).
fof(f482, plain, ((e3 = op(e3, e1)) | ~ spl30_12), inference(avatar_component_clause, [], [f480])).
fof(f5426, plain, (~ spl30_1 | ~ spl30_13), inference(avatar_split_clause, [], [f5421, f485, f434])).
fof(f434, plain, (spl30_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_1])])).
fof(f5421, plain, (~ (e0 = op(e3, e3)) | ~ spl30_13), inference(backward_demodulation, [], [f207, f487])).
fof(f487, plain, ((e0 = op(e3, e0)) | ~ spl30_13), inference(avatar_component_clause, [], [f485])).
fof(f5418, plain, (~ spl30_1 | ~ spl30_17), inference(avatar_split_clause, [], [f5413, f502, f434])).
fof(f5413, plain, (~ (e0 = op(e3, e3)) | ~ spl30_17), inference(backward_demodulation, [], [f186, f504])).
fof(f504, plain, ((e0 = op(e2, e3)) | ~ spl30_17), inference(avatar_component_clause, [], [f502])).
fof(f5412, plain, (~ spl30_18 | ~ spl30_26), inference(avatar_split_clause, [], [f5407, f540, f506])).
fof(f506, plain, (spl30_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_18])])).
fof(f540, plain, (spl30_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_26])])).
fof(f5407, plain, (~ (e1 = op(e2, e3)) | ~ spl30_26), inference(backward_demodulation, [], [f203, f542])).
fof(f542, plain, ((e1 = op(e2, e1)) | ~ spl30_26), inference(avatar_component_clause, [], [f540])).
fof(f203, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5411, plain, (~ spl30_10 | ~ spl30_26), inference(avatar_split_clause, [], [f5406, f540, f472])).
fof(f5406, plain, (~ (e1 = op(e3, e1)) | ~ spl30_26), inference(backward_demodulation, [], [f174, f542])).
fof(f5405, plain, (spl30_17 | ~ spl30_32 | ~ spl30_69), inference(avatar_split_clause, [], [f5403, f757, f565, f502])).
fof(f757, plain, (spl30_69 <=> (e0 = op(e2, op(e2, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_69])])).
fof(f5403, plain, ((e0 = op(e2, e3)) | (~ spl30_32 | ~ spl30_69)), inference(backward_demodulation, [], [f758, f567])).
fof(f758, plain, ((e0 = op(e2, op(e2, e0))) | ~ spl30_69), inference(avatar_component_clause, [], [f757])).
fof(f5404, plain, (~ spl30_28 | ~ spl30_32), inference(avatar_split_clause, [], [f5401, f565, f548])).
fof(f5401, plain, (~ (e3 = op(e2, e1)) | ~ spl30_32), inference(backward_demodulation, [], [f199, f567])).
fof(f199, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f5398, plain, (~ spl30_18 | ~ spl30_34), inference(avatar_split_clause, [], [f5393, f574, f506])).
fof(f5393, plain, (~ (e1 = op(e2, e3)) | ~ spl30_34), inference(backward_demodulation, [], [f184, f576])).
fof(f576, plain, ((e1 = op(e1, e3)) | ~ spl30_34), inference(avatar_component_clause, [], [f574])).
fof(f5392, plain, (~ spl30_36 | ~ spl30_40), inference(avatar_split_clause, [], [f5389, f599, f582])).
fof(f582, plain, (spl30_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_36])])).
fof(f5389, plain, (~ (e3 = op(e1, e3)) | ~ spl30_40), inference(backward_demodulation, [], [f198, f601])).
fof(f198, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f5391, plain, (~ spl30_8 | ~ spl30_40), inference(avatar_split_clause, [], [f5388, f599, f463])).
fof(f5388, plain, (~ (e3 = op(e3, e2)) | ~ spl30_40), inference(backward_demodulation, [], [f179, f601])).
fof(f179, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f5386, plain, (~ spl30_15 | ~ spl30_41 | ~ spl30_119 | spl30_120), inference(avatar_split_clause, [], [f5376, f1158, f1153, f604, f493])).
fof(f1158, plain, (spl30_120 <=> (e2 = op(op(op(e1, e1), e1), op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_120])])).
fof(f5376, plain, (~ (e2 = op(e3, e0)) | (~ spl30_41 | ~ spl30_119 | spl30_120)), inference(backward_demodulation, [], [f5303, f606])).
fof(f5303, plain, (~ (e2 = op(e3, op(e1, e1))) | (~ spl30_119 | spl30_120)), inference(forward_demodulation, [], [f1160, f1154])).
fof(f1160, plain, (~ (e2 = op(op(op(e1, e1), e1), op(e1, e1))) | spl30_120), inference(avatar_component_clause, [], [f1158])).
fof(f5380, plain, (~ spl30_41 | ~ spl30_59 | spl30_118), inference(avatar_contradiction_clause, [], [f5379])).
fof(f5379, plain, ($false | (~ spl30_41 | ~ spl30_59 | spl30_118)), inference(subsumption_resolution, [], [f5372, f682])).
fof(f5372, plain, (~ (e2 = op(e0, e1)) | (~ spl30_41 | spl30_118)), inference(backward_demodulation, [], [f1150, f606])).
fof(f1150, plain, (~ (e2 = op(op(e1, e1), e1)) | spl30_118), inference(avatar_component_clause, [], [f1148])).
fof(f1148, plain, (spl30_118 <=> (e2 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_118])])).
fof(f5378, plain, (~ spl30_37 | ~ spl30_41), inference(avatar_split_clause, [], [f5366, f604, f587])).
fof(f587, plain, (spl30_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_37])])).
fof(f5366, plain, (~ (e0 = op(e1, e2)) | ~ spl30_41), inference(backward_demodulation, [], [f196, f606])).
fof(f196, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f5363, plain, (~ spl30_37 | ~ spl30_47 | spl30_74), inference(avatar_split_clause, [], [f5360, f783, f629, f587])).
fof(f629, plain, (spl30_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_47])])).
fof(f5360, plain, (~ (e0 = op(e1, e2)) | (~ spl30_47 | spl30_74)), inference(backward_demodulation, [], [f785, f631])).
fof(f631, plain, ((e2 = op(e1, e0)) | ~ spl30_47), inference(avatar_component_clause, [], [f629])).
fof(f5362, plain, (~ spl30_43 | ~ spl30_47), inference(avatar_split_clause, [], [f5356, f629, f612])).
fof(f5356, plain, (~ (e2 = op(e1, e1)) | ~ spl30_47), inference(backward_demodulation, [], [f193, f631])).
fof(f193, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f5361, plain, (~ spl30_15 | ~ spl30_47), inference(avatar_split_clause, [], [f5355, f629, f493])).
fof(f5355, plain, (~ (e2 = op(e3, e0)) | ~ spl30_47), inference(backward_demodulation, [], [f167, f631])).
fof(f167, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f5353, plain, (~ spl30_36 | ~ spl30_52), inference(avatar_split_clause, [], [f5349, f650, f582])).
fof(f5349, plain, (~ (e3 = op(e1, e3)) | ~ spl30_52), inference(backward_demodulation, [], [f181, f652])).
fof(f652, plain, ((e3 = op(e0, e3)) | ~ spl30_52), inference(avatar_component_clause, [], [f650])).
fof(f5347, plain, (~ spl30_37 | ~ spl30_53), inference(avatar_split_clause, [], [f5340, f655, f587])).
fof(f5340, plain, (~ (e0 = op(e1, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f175, f657])).
fof(f657, plain, ((e0 = op(e0, e2)) | ~ spl30_53), inference(avatar_component_clause, [], [f655])).
fof(f175, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f5339, plain, (~ spl30_54 | ~ spl30_59 | spl30_78), inference(avatar_split_clause, [], [f5336, f804, f680, f659])).
fof(f659, plain, (spl30_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_54])])).
fof(f804, plain, (spl30_78 <=> (e1 = op(e0, op(e0, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_78])])).
fof(f5336, plain, (~ (e1 = op(e0, e2)) | (~ spl30_59 | spl30_78)), inference(backward_demodulation, [], [f806, f682])).
fof(f806, plain, (~ (e1 = op(e0, op(e0, e1))) | spl30_78), inference(avatar_component_clause, [], [f804])).
fof(f5338, plain, (~ spl30_51 | ~ spl30_59), inference(avatar_split_clause, [], [f5334, f680, f646])).
fof(f646, plain, (spl30_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_51])])).
fof(f5334, plain, (~ (e2 = op(e0, e3)) | ~ spl30_59), inference(backward_demodulation, [], [f191, f682])).
fof(f191, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f5337, plain, (~ spl30_43 | ~ spl30_59), inference(avatar_split_clause, [], [f5330, f680, f612])).
fof(f5330, plain, (~ (e2 = op(e1, e1)) | ~ spl30_59), inference(backward_demodulation, [], [f169, f682])).
fof(f5326, plain, (~ spl30_57 | ~ spl30_62 | spl30_79), inference(avatar_split_clause, [], [f5319, f809, f693, f672])).
fof(f693, plain, (spl30_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl30_62])])).
fof(f809, plain, (spl30_79 <=> (e0 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_79])])).
fof(f5319, plain, (~ (e0 = op(e0, e1)) | (~ spl30_62 | spl30_79)), inference(backward_demodulation, [], [f811, f695])).
fof(f695, plain, ((op(e0, e0) = e1) | ~ spl30_62), inference(avatar_component_clause, [], [f693])).
fof(f811, plain, (~ (e0 = op(e0, op(e0, e0))) | spl30_79), inference(avatar_component_clause, [], [f809])).
fof(f5325, plain, (~ spl30_54 | ~ spl30_62), inference(avatar_split_clause, [], [f5315, f693, f659])).
fof(f5315, plain, (~ (e1 = op(e0, e2)) | ~ spl30_62), inference(backward_demodulation, [], [f188, f695])).
fof(f188, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f5324, plain, (~ spl30_46 | ~ spl30_62), inference(avatar_split_clause, [], [f5311, f693, f625])).
fof(f5311, plain, (~ (e1 = op(e1, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f163, f695])).
fof(f163, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f5309, plain, (~ spl30_131 | ~ spl30_133), inference(avatar_contradiction_clause, [], [f5308])).
fof(f5308, plain, ($false | (~ spl30_131 | ~ spl30_133)), inference(subsumption_resolution, [], [f5306, f212])).
fof(f212, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax4)).
fof(f5306, plain, ((e0 = e2) | (~ spl30_131 | ~ spl30_133)), inference(backward_demodulation, [], [f1214, f1223])).
fof(f1223, plain, ((e0 = op(op(e3, e3), e3)) | ~ spl30_133), inference(avatar_component_clause, [], [f1222])).
fof(f1222, plain, (spl30_133 <=> (e0 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_133])])).
fof(f1214, plain, ((e2 = op(op(e3, e3), e3)) | ~ spl30_131), inference(avatar_component_clause, [], [f1213])).
fof(f5288, plain, (~ spl30_19 | ~ spl30_23), inference(avatar_split_clause, [], [f5282, f527, f510])).
fof(f5282, plain, (~ (e2 = op(e2, e3)) | ~ spl30_23), inference(backward_demodulation, [], [f204, f529])).
fof(f529, plain, ((e2 = op(e2, e2)) | ~ spl30_23), inference(avatar_component_clause, [], [f527])).
fof(f204, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5281, plain, (spl30_18 | ~ spl30_28 | ~ spl30_68), inference(avatar_split_clause, [], [f5277, f752, f548, f506])).
fof(f752, plain, (spl30_68 <=> (e1 = op(e2, op(e2, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_68])])).
fof(f5277, plain, ((e1 = op(e2, e3)) | (~ spl30_28 | ~ spl30_68)), inference(backward_demodulation, [], [f753, f550])).
fof(f753, plain, ((e1 = op(e2, op(e2, e1))) | ~ spl30_68), inference(avatar_component_clause, [], [f752])).
fof(f5274, plain, (~ spl30_25 | ~ spl30_29), inference(avatar_split_clause, [], [f5267, f553, f536])).
fof(f536, plain, (spl30_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_25])])).
fof(f5267, plain, (~ (e0 = op(e2, e1)) | ~ spl30_29), inference(backward_demodulation, [], [f199, f555])).
fof(f5264, plain, (~ spl30_19 | ~ spl30_51), inference(avatar_split_clause, [], [f5260, f646, f510])).
fof(f5260, plain, (~ (e2 = op(e2, e3)) | ~ spl30_51), inference(backward_demodulation, [], [f182, f648])).
fof(f648, plain, ((e2 = op(e0, e3)) | ~ spl30_51), inference(avatar_component_clause, [], [f646])).
fof(f182, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f5258, plain, (~ spl30_50 | ~ spl30_54), inference(avatar_split_clause, [], [f5253, f659, f642])).
fof(f5253, plain, (~ (e1 = op(e0, e3)) | ~ spl30_54), inference(backward_demodulation, [], [f192, f661])).
fof(f661, plain, ((e1 = op(e0, e2)) | ~ spl30_54), inference(avatar_component_clause, [], [f659])).
fof(f192, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f5257, plain, (~ spl30_22 | ~ spl30_54), inference(avatar_split_clause, [], [f5252, f659, f523])).
fof(f5252, plain, (~ (e1 = op(e2, e2)) | ~ spl30_54), inference(backward_demodulation, [], [f176, f661])).
fof(f5251, plain, (~ spl30_25 | ~ spl30_57), inference(avatar_split_clause, [], [f5244, f672, f536])).
fof(f5244, plain, (~ (e0 = op(e2, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f170, f674])).
fof(f170, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f5243, plain, (~ spl30_18 | ~ spl30_64 | spl30_111 | ~ spl30_112), inference(avatar_split_clause, [], [f5239, f1118, f1114, f701, f506])).
fof(f1114, plain, (spl30_111 <=> (e1 = op(op(op(e0, e0), e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_111])])).
fof(f1118, plain, (spl30_112 <=> (e2 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_112])])).
fof(f5239, plain, (~ (e1 = op(e2, e3)) | (~ spl30_64 | spl30_111 | ~ spl30_112)), inference(backward_demodulation, [], [f5229, f703])).
fof(f5229, plain, (~ (e1 = op(e2, op(e0, e0))) | (spl30_111 | ~ spl30_112)), inference(backward_demodulation, [], [f1116, f1119])).
fof(f1119, plain, ((e2 = op(op(e0, e0), e0)) | ~ spl30_112), inference(avatar_component_clause, [], [f1118])).
fof(f1116, plain, (~ (e1 = op(op(op(e0, e0), e0), op(e0, e0))) | spl30_111), inference(avatar_component_clause, [], [f1114])).
fof(f5242, plain, (~ spl30_60 | ~ spl30_64), inference(avatar_split_clause, [], [f5233, f701, f684])).
fof(f684, plain, (spl30_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_60])])).
fof(f5233, plain, (~ (e3 = op(e0, e1)) | ~ spl30_64), inference(backward_demodulation, [], [f187, f703])).
fof(f187, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f5226, plain, (~ spl30_25 | ~ spl30_43 | spl30_121), inference(avatar_split_clause, [], [f5225, f1162, f612, f536])).
fof(f1162, plain, (spl30_121 <=> (e0 = op(op(e1, e1), e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_121])])).
fof(f5225, plain, (~ (e0 = op(e2, e1)) | (~ spl30_43 | spl30_121)), inference(forward_demodulation, [], [f1164, f614])).
fof(f1164, plain, (~ (e0 = op(op(e1, e1), e1)) | spl30_121), inference(avatar_component_clause, [], [f1162])).
fof(f5220, plain, (~ spl30_50 | ~ spl30_1 | spl30_130), inference(avatar_split_clause, [], [f5219, f1208, f434, f642])).
fof(f1208, plain, (spl30_130 <=> (e1 = op(op(e3, e3), e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_130])])).
fof(f5219, plain, (~ (e1 = op(e0, e3)) | (~ spl30_1 | spl30_130)), inference(forward_demodulation, [], [f1210, f436])).
fof(f436, plain, ((e0 = op(e3, e3)) | ~ spl30_1), inference(avatar_component_clause, [], [f434])).
fof(f1210, plain, (~ (e1 = op(op(e3, e3), e3)) | spl30_130), inference(avatar_component_clause, [], [f1208])).
fof(f5218, plain, (spl30_51 | ~ spl30_1 | ~ spl30_131), inference(avatar_split_clause, [], [f5217, f1213, f434, f646])).
fof(f5217, plain, ((e2 = op(e0, e3)) | (~ spl30_1 | ~ spl30_131)), inference(forward_demodulation, [], [f1214, f436])).
fof(f5172, plain, (~ spl30_22 | ~ spl30_60 | ~ spl30_127 | spl30_128), inference(avatar_contradiction_clause, [], [f5171])).
fof(f5171, plain, ($false | (~ spl30_22 | ~ spl30_60 | ~ spl30_127 | spl30_128)), inference(subsumption_resolution, [], [f5166, f686])).
fof(f686, plain, ((e3 = op(e0, e1)) | ~ spl30_60), inference(avatar_component_clause, [], [f684])).
fof(f5166, plain, (~ (e3 = op(e0, e1)) | (~ spl30_22 | ~ spl30_127 | spl30_128)), inference(backward_demodulation, [], [f5077, f525])).
fof(f5077, plain, (~ (e3 = op(e0, op(e2, e2))) | (~ spl30_127 | spl30_128)), inference(backward_demodulation, [], [f1200, f1193])).
fof(f1193, plain, ((e0 = op(op(e2, e2), e2)) | ~ spl30_127), inference(avatar_component_clause, [], [f1192])).
fof(f1192, plain, (spl30_127 <=> (e0 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_127])])).
fof(f1200, plain, (~ (e3 = op(op(op(e2, e2), e2), op(e2, e2))) | spl30_128), inference(avatar_component_clause, [], [f1198])).
fof(f1198, plain, (spl30_128 <=> (e3 = op(op(op(e2, e2), e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_128])])).
fof(f5169, plain, (~ spl30_18 | ~ spl30_22), inference(avatar_split_clause, [], [f5161, f523, f506])).
fof(f5161, plain, (~ (e1 = op(e2, e3)) | ~ spl30_22), inference(backward_demodulation, [], [f204, f525])).
fof(f5159, plain, (~ spl30_9 | ~ spl30_25), inference(avatar_split_clause, [], [f5152, f536, f468])).
fof(f468, plain, (spl30_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_9])])).
fof(f5152, plain, (~ (e0 = op(e3, e1)) | ~ spl30_25), inference(backward_demodulation, [], [f174, f538])).
fof(f538, plain, ((e0 = op(e2, e1)) | ~ spl30_25), inference(avatar_component_clause, [], [f536])).
fof(f5151, plain, (~ spl30_24 | ~ spl30_32), inference(avatar_split_clause, [], [f5147, f565, f531])).
fof(f5147, plain, (~ (e3 = op(e2, e2)) | ~ spl30_32), inference(backward_demodulation, [], [f200, f567])).
fof(f200, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f5150, plain, (~ spl30_16 | ~ spl30_32), inference(avatar_split_clause, [], [f5145, f565, f497])).
fof(f497, plain, (spl30_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_16])])).
fof(f5145, plain, (~ (e3 = op(e3, e0)) | ~ spl30_32), inference(backward_demodulation, [], [f168, f567])).
fof(f5134, plain, (spl30_25 | ~ spl30_43 | ~ spl30_121), inference(avatar_split_clause, [], [f5128, f1162, f612, f536])).
fof(f5128, plain, ((e0 = op(e2, e1)) | (~ spl30_43 | ~ spl30_121)), inference(backward_demodulation, [], [f1163, f614])).
fof(f1163, plain, ((e0 = op(op(e1, e1), e1)) | ~ spl30_121), inference(avatar_component_clause, [], [f1162])).
fof(f5133, plain, (~ spl30_39 | ~ spl30_43), inference(avatar_split_clause, [], [f5125, f612, f595])).
fof(f595, plain, (spl30_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_39])])).
fof(f5125, plain, (~ (e2 = op(e1, e2)) | ~ spl30_43), inference(backward_demodulation, [], [f196, f614])).
fof(f5132, plain, (~ spl30_27 | ~ spl30_43), inference(avatar_split_clause, [], [f5123, f612, f544])).
fof(f544, plain, (spl30_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_27])])).
fof(f5123, plain, (~ (e2 = op(e2, e1)) | ~ spl30_43), inference(backward_demodulation, [], [f172, f614])).
fof(f172, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f5122, plain, (~ spl30_18 | ~ spl30_50), inference(avatar_split_clause, [], [f5117, f642, f506])).
fof(f5117, plain, (~ (e1 = op(e2, e3)) | ~ spl30_50), inference(backward_demodulation, [], [f182, f644])).
fof(f5116, plain, (~ spl30_39 | ~ spl30_55), inference(avatar_split_clause, [], [f5110, f663, f595])).
fof(f5110, plain, (~ (e2 = op(e1, e2)) | ~ spl30_55), inference(backward_demodulation, [], [f175, f665])).
fof(f5102, plain, (~ spl30_49 | ~ spl30_61), inference(avatar_split_clause, [], [f5090, f689, f638])).
fof(f638, plain, (spl30_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_49])])).
fof(f689, plain, (spl30_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_61])])).
fof(f5090, plain, (~ (e0 = op(e0, e3)) | ~ spl30_61), inference(backward_demodulation, [], [f189, f691])).
fof(f691, plain, ((e0 = op(e0, e0)) | ~ spl30_61), inference(avatar_component_clause, [], [f689])).
fof(f5101, plain, (~ spl30_53 | ~ spl30_61), inference(avatar_split_clause, [], [f5089, f689, f655])).
fof(f5089, plain, (~ (e0 = op(e0, e2)) | ~ spl30_61), inference(backward_demodulation, [], [f188, f691])).
fof(f5100, plain, (~ spl30_29 | ~ spl30_61), inference(avatar_split_clause, [], [f5087, f689, f553])).
fof(f5087, plain, (~ (e0 = op(e2, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f164, f691])).
fof(f5073, plain, (~ spl30_30 | ~ spl30_46), inference(avatar_split_clause, [], [f5072, f625, f557])).
fof(f557, plain, (spl30_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_30])])).
fof(f5072, plain, (~ (e1 = op(e2, e0)) | ~ spl30_46), inference(forward_demodulation, [], [f166, f627])).
fof(f166, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f5052, plain, (spl30_41 | ~ spl30_46 | ~ spl30_74), inference(avatar_split_clause, [], [f5050, f783, f625, f604])).
fof(f5050, plain, ((e0 = op(e1, e1)) | (~ spl30_46 | ~ spl30_74)), inference(backward_demodulation, [], [f784, f627])).
fof(f784, plain, ((e0 = op(e1, op(e1, e0))) | ~ spl30_74), inference(avatar_component_clause, [], [f783])).
fof(f5051, plain, (~ spl30_42 | ~ spl30_46), inference(avatar_split_clause, [], [f5049, f625, f608])).
fof(f608, plain, (spl30_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_42])])).
fof(f5049, plain, (~ (e1 = op(e1, e1)) | ~ spl30_46), inference(backward_demodulation, [], [f193, f627])).
fof(f5039, plain, (~ spl30_58 | ~ spl30_60), inference(avatar_contradiction_clause, [], [f5038])).
fof(f5038, plain, ($false | (~ spl30_58 | ~ spl30_60)), inference(subsumption_resolution, [], [f5037, f215])).
fof(f215, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f5037, plain, ((e1 = e3) | (~ spl30_58 | ~ spl30_60)), inference(backward_demodulation, [], [f686, f678])).
fof(f678, plain, ((e1 = op(e0, e1)) | ~ spl30_58), inference(avatar_component_clause, [], [f676])).
fof(f676, plain, (spl30_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_58])])).
fof(f5030, plain, (spl30_53 | ~ spl30_63 | ~ spl30_79), inference(avatar_split_clause, [], [f5023, f809, f697, f655])).
fof(f697, plain, (spl30_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl30_63])])).
fof(f5023, plain, ((e0 = op(e0, e2)) | (~ spl30_63 | ~ spl30_79)), inference(backward_demodulation, [], [f810, f699])).
fof(f699, plain, ((op(e0, e0) = e2) | ~ spl30_63), inference(avatar_component_clause, [], [f697])).
fof(f810, plain, ((e0 = op(e0, op(e0, e0))) | ~ spl30_79), inference(avatar_component_clause, [], [f809])).
fof(f5029, plain, (~ spl30_55 | ~ spl30_63), inference(avatar_split_clause, [], [f5021, f697, f663])).
fof(f5021, plain, (~ (e2 = op(e0, e2)) | ~ spl30_63), inference(backward_demodulation, [], [f188, f699])).
fof(f5028, plain, (~ spl30_47 | ~ spl30_63), inference(avatar_split_clause, [], [f5020, f697, f629])).
fof(f5020, plain, (~ (e2 = op(e1, e0)) | ~ spl30_63), inference(backward_demodulation, [], [f163, f699])).
fof(f5001, plain, (~ spl30_53 | ~ spl30_49), inference(avatar_split_clause, [], [f5000, f638, f655])).
fof(f5000, plain, (~ (e0 = op(e0, e2)) | ~ spl30_49), inference(forward_demodulation, [], [f192, f640])).
fof(f640, plain, ((e0 = op(e0, e3)) | ~ spl30_49), inference(avatar_component_clause, [], [f638])).
fof(f4999, plain, (~ spl30_56 | ~ spl30_60), inference(avatar_split_clause, [], [f4998, f684, f667])).
fof(f667, plain, (spl30_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_56])])).
fof(f4998, plain, (~ (e3 = op(e0, e2)) | ~ spl30_60), inference(forward_demodulation, [], [f190, f686])).
fof(f4996, plain, (~ spl30_48 | ~ spl30_16), inference(avatar_split_clause, [], [f4995, f497, f633])).
fof(f633, plain, (spl30_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_48])])).
fof(f4995, plain, (~ (e3 = op(e1, e0)) | ~ spl30_16), inference(forward_demodulation, [], [f167, f499])).
fof(f499, plain, ((e3 = op(e3, e0)) | ~ spl30_16), inference(avatar_component_clause, [], [f497])).
fof(f4992, plain, (~ spl30_41 | ~ spl30_9), inference(avatar_split_clause, [], [f4991, f468, f604])).
fof(f4991, plain, (~ (e0 = op(e1, e1)) | ~ spl30_9), inference(forward_demodulation, [], [f173, f470])).
fof(f470, plain, ((e0 = op(e3, e1)) | ~ spl30_9), inference(avatar_component_clause, [], [f468])).
fof(f173, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4989, plain, (~ spl30_44 | ~ spl30_60), inference(avatar_split_clause, [], [f4988, f684, f616])).
fof(f616, plain, (spl30_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_44])])).
fof(f4988, plain, (~ (e3 = op(e1, e1)) | ~ spl30_60), inference(forward_demodulation, [], [f169, f686])).
fof(f4984, plain, (~ spl30_37 | ~ spl30_3 | spl30_129 | ~ spl30_130), inference(avatar_split_clause, [], [f4977, f1208, f1204, f442, f587])).
fof(f1204, plain, (spl30_129 <=> (e0 = op(op(op(e3, e3), e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_129])])).
fof(f4977, plain, (~ (e0 = op(e1, e2)) | (~ spl30_3 | spl30_129 | ~ spl30_130)), inference(backward_demodulation, [], [f4766, f444])).
fof(f444, plain, ((e2 = op(e3, e3)) | ~ spl30_3), inference(avatar_component_clause, [], [f442])).
fof(f4766, plain, (~ (e0 = op(e1, op(e3, e3))) | (spl30_129 | ~ spl30_130)), inference(backward_demodulation, [], [f1206, f1209])).
fof(f1209, plain, ((e1 = op(op(e3, e3), e3)) | ~ spl30_130), inference(avatar_component_clause, [], [f1208])).
fof(f1206, plain, (~ (e0 = op(op(op(e3, e3), e3), op(e3, e3))) | spl30_129), inference(avatar_component_clause, [], [f1204])).
fof(f4967, plain, (~ spl30_1 | ~ spl30_9), inference(avatar_split_clause, [], [f4962, f468, f434])).
fof(f4962, plain, (~ (e0 = op(e3, e3)) | ~ spl30_9), inference(backward_demodulation, [], [f209, f470])).
fof(f209, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4960, plain, (~ spl30_1 | ~ spl30_16 | spl30_84), inference(avatar_split_clause, [], [f4958, f835, f497, f434])).
fof(f4958, plain, (~ (e0 = op(e3, e3)) | (~ spl30_16 | spl30_84)), inference(backward_demodulation, [], [f837, f499])).
fof(f837, plain, (~ (e0 = op(e3, op(e3, e0))) | spl30_84), inference(avatar_component_clause, [], [f835])).
fof(f4959, plain, (~ spl30_8 | ~ spl30_16), inference(avatar_split_clause, [], [f4956, f497, f463])).
fof(f4956, plain, (~ (e3 = op(e3, e2)) | ~ spl30_16), inference(backward_demodulation, [], [f206, f499])).
fof(f206, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4948, plain, (~ spl30_19 | ~ spl30_24 | spl30_67), inference(avatar_split_clause, [], [f4942, f747, f531, f510])).
fof(f747, plain, (spl30_67 <=> (e2 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_67])])).
fof(f4942, plain, (~ (e2 = op(e2, e3)) | (~ spl30_24 | spl30_67)), inference(backward_demodulation, [], [f749, f533])).
fof(f533, plain, ((e3 = op(e2, e2)) | ~ spl30_24), inference(avatar_component_clause, [], [f531])).
fof(f749, plain, (~ (e2 = op(e2, op(e2, e2))) | spl30_67), inference(avatar_component_clause, [], [f747])).
fof(f4947, plain, (~ spl30_8 | ~ spl30_24), inference(avatar_split_clause, [], [f4940, f531, f463])).
fof(f4940, plain, (~ (e3 = op(e3, e2)) | ~ spl30_24), inference(backward_demodulation, [], [f180, f533])).
fof(f4939, plain, (~ spl30_22 | ~ spl30_27 | spl30_68), inference(avatar_split_clause, [], [f4936, f752, f544, f523])).
fof(f4936, plain, (~ (e1 = op(e2, e2)) | (~ spl30_27 | spl30_68)), inference(backward_demodulation, [], [f754, f546])).
fof(f546, plain, ((e2 = op(e2, e1)) | ~ spl30_27), inference(avatar_component_clause, [], [f544])).
fof(f754, plain, (~ (e1 = op(e2, op(e2, e1))) | spl30_68), inference(avatar_component_clause, [], [f752])).
fof(f4938, plain, (~ spl30_19 | ~ spl30_27), inference(avatar_split_clause, [], [f4934, f544, f510])).
fof(f4934, plain, (~ (e2 = op(e2, e3)) | ~ spl30_27), inference(backward_demodulation, [], [f203, f546])).
fof(f4937, plain, (~ spl30_11 | ~ spl30_27), inference(avatar_split_clause, [], [f4932, f544, f476])).
fof(f476, plain, (spl30_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl30_11])])).
fof(f4932, plain, (~ (e2 = op(e3, e1)) | ~ spl30_27), inference(backward_demodulation, [], [f174, f546])).
fof(f4931, plain, (~ spl30_1 | ~ spl30_49), inference(avatar_split_clause, [], [f4928, f638, f434])).
fof(f4928, plain, (~ (e0 = op(e3, e3)) | ~ spl30_49), inference(backward_demodulation, [], [f183, f640])).
fof(f183, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4925, plain, (~ spl30_50 | ~ spl30_60 | spl30_78), inference(avatar_split_clause, [], [f4923, f804, f684, f642])).
fof(f4923, plain, (~ (e1 = op(e0, e3)) | (~ spl30_60 | spl30_78)), inference(backward_demodulation, [], [f806, f686])).
fof(f4924, plain, (~ spl30_28 | ~ spl30_60), inference(avatar_split_clause, [], [f4920, f684, f548])).
fof(f4920, plain, (~ (e3 = op(e2, e1)) | ~ spl30_60), inference(backward_demodulation, [], [f170, f686])).
fof(f4917, plain, (~ spl30_50 | ~ spl30_62), inference(avatar_split_clause, [], [f4910, f693, f642])).
fof(f4910, plain, (~ (e1 = op(e0, e3)) | ~ spl30_62), inference(backward_demodulation, [], [f189, f695])).
fof(f4916, plain, (~ spl30_14 | ~ spl30_62), inference(avatar_split_clause, [], [f4908, f693, f489])).
fof(f489, plain, (spl30_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_14])])).
fof(f4908, plain, (~ (e1 = op(e3, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f165, f695])).
fof(f165, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f4895, plain, (~ spl30_1 | ~ spl30_47 | ~ spl30_130 | spl30_134), inference(avatar_contradiction_clause, [], [f4894])).
fof(f4894, plain, ($false | (~ spl30_1 | ~ spl30_47 | ~ spl30_130 | spl30_134)), inference(subsumption_resolution, [], [f4892, f631])).
fof(f4892, plain, (~ (e2 = op(e1, e0)) | (~ spl30_1 | ~ spl30_130 | spl30_134)), inference(backward_demodulation, [], [f4767, f436])).
fof(f4767, plain, (~ (e2 = op(e1, op(e3, e3))) | (~ spl30_130 | spl30_134)), inference(backward_demodulation, [], [f1230, f1209])).
fof(f1230, plain, (~ (e2 = op(op(op(e3, e3), e3), op(e3, e3))) | spl30_134), inference(avatar_component_clause, [], [f1228])).
fof(f1228, plain, (spl30_134 <=> (e2 = op(op(op(e3, e3), e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_134])])).
fof(f4882, plain, (~ spl30_7 | ~ spl30_11), inference(avatar_split_clause, [], [f4878, f476, f459])).
fof(f459, plain, (spl30_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_7])])).
fof(f4878, plain, (~ (e2 = op(e3, e2)) | ~ spl30_11), inference(backward_demodulation, [], [f208, f478])).
fof(f478, plain, ((e2 = op(e3, e1)) | ~ spl30_11), inference(avatar_component_clause, [], [f476])).
fof(f4877, plain, (~ spl30_9 | ~ spl30_14 | spl30_84), inference(avatar_split_clause, [], [f4875, f835, f489, f468])).
fof(f4875, plain, (~ (e0 = op(e3, e1)) | (~ spl30_14 | spl30_84)), inference(backward_demodulation, [], [f837, f491])).
fof(f491, plain, ((e1 = op(e3, e0)) | ~ spl30_14), inference(avatar_component_clause, [], [f489])).
fof(f4876, plain, (~ spl30_2 | ~ spl30_14), inference(avatar_split_clause, [], [f4872, f489, f438])).
fof(f438, plain, (spl30_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_2])])).
fof(f4872, plain, (~ (e1 = op(e3, e3)) | ~ spl30_14), inference(backward_demodulation, [], [f207, f491])).
fof(f4858, plain, (~ spl30_20 | ~ spl30_28), inference(avatar_split_clause, [], [f4856, f548, f514])).
fof(f514, plain, (spl30_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_20])])).
fof(f4856, plain, (~ (e3 = op(e2, e3)) | ~ spl30_28), inference(backward_demodulation, [], [f203, f550])).
fof(f4853, plain, (~ spl30_21 | ~ spl30_29), inference(avatar_split_clause, [], [f4847, f553, f519])).
fof(f519, plain, (spl30_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_21])])).
fof(f4847, plain, (~ (e0 = op(e2, e2)) | ~ spl30_29), inference(backward_demodulation, [], [f200, f555])).
fof(f4844, plain, (~ spl30_20 | ~ spl30_36), inference(avatar_split_clause, [], [f4841, f582, f514])).
fof(f4841, plain, (~ (e3 = op(e2, e3)) | ~ spl30_36), inference(backward_demodulation, [], [f184, f584])).
fof(f584, plain, ((e3 = op(e1, e3)) | ~ spl30_36), inference(avatar_component_clause, [], [f582])).
fof(f4840, plain, (~ spl30_33 | ~ spl30_37), inference(avatar_split_clause, [], [f4834, f587, f570])).
fof(f4834, plain, (~ (e0 = op(e1, e3)) | ~ spl30_37), inference(backward_demodulation, [], [f198, f589])).
fof(f589, plain, ((e0 = op(e1, e2)) | ~ spl30_37), inference(avatar_component_clause, [], [f587])).
fof(f4839, plain, (~ spl30_21 | ~ spl30_37), inference(avatar_split_clause, [], [f4832, f587, f519])).
fof(f4832, plain, (~ (e0 = op(e2, e2)) | ~ spl30_37), inference(backward_demodulation, [], [f178, f589])).
fof(f4827, plain, (~ spl30_38 | ~ spl30_42), inference(avatar_split_clause, [], [f4818, f608, f591])).
fof(f591, plain, (spl30_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_38])])).
fof(f4818, plain, (~ (e1 = op(e1, e2)) | ~ spl30_42), inference(backward_demodulation, [], [f196, f610])).
fof(f610, plain, ((e1 = op(e1, e1)) | ~ spl30_42), inference(avatar_component_clause, [], [f608])).
fof(f4815, plain, (~ spl30_2 | ~ spl30_50), inference(avatar_split_clause, [], [f4812, f642, f438])).
fof(f4812, plain, (~ (e1 = op(e3, e3)) | ~ spl30_50), inference(backward_demodulation, [], [f183, f644])).
fof(f4809, plain, (~ spl30_51 | ~ spl30_55), inference(avatar_split_clause, [], [f4805, f663, f646])).
fof(f4805, plain, (~ (e2 = op(e0, e3)) | ~ spl30_55), inference(backward_demodulation, [], [f192, f665])).
fof(f4808, plain, (~ spl30_7 | ~ spl30_55), inference(avatar_split_clause, [], [f4804, f663, f459])).
fof(f4804, plain, (~ (e2 = op(e3, e2)) | ~ spl30_55), inference(backward_demodulation, [], [f177, f665])).
fof(f177, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f4801, plain, (~ spl30_9 | ~ spl30_57), inference(avatar_split_clause, [], [f4794, f672, f468])).
fof(f4794, plain, (~ (e0 = op(e3, e1)) | ~ spl30_57), inference(backward_demodulation, [], [f171, f674])).
fof(f171, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4789, plain, (~ spl30_56 | ~ spl30_64), inference(avatar_split_clause, [], [f4782, f701, f667])).
fof(f4782, plain, (~ (e3 = op(e0, e2)) | ~ spl30_64), inference(backward_demodulation, [], [f188, f703])).
fof(f4788, plain, (~ spl30_16 | ~ spl30_64), inference(avatar_split_clause, [], [f4780, f701, f497])).
fof(f4780, plain, (~ (e3 = op(e3, e0)) | ~ spl30_64), inference(backward_demodulation, [], [f165, f703])).
fof(f4778, plain, (spl30_37 | ~ spl30_47 | ~ spl30_74), inference(avatar_split_clause, [], [f4777, f783, f629, f587])).
fof(f4777, plain, ((e0 = op(e1, e2)) | (~ spl30_47 | ~ spl30_74)), inference(forward_demodulation, [], [f784, f631])).
fof(f4763, plain, (~ spl30_39 | ~ spl30_47), inference(avatar_split_clause, [], [f4762, f629, f595])).
fof(f4762, plain, (~ (e2 = op(e1, e2)) | ~ spl30_47), inference(forward_demodulation, [], [f194, f631])).
fof(f194, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f4761, plain, (~ spl30_31 | ~ spl30_47), inference(avatar_split_clause, [], [f4760, f629, f561])).
fof(f561, plain, (spl30_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_31])])).
fof(f4760, plain, (~ (e2 = op(e2, e0)) | ~ spl30_47), inference(forward_demodulation, [], [f166, f631])).
fof(f4756, plain, (~ spl30_9 | ~ spl30_44 | ~ spl30_51 | spl30_120), inference(avatar_contradiction_clause, [], [f4755])).
fof(f4755, plain, ($false | (~ spl30_9 | ~ spl30_44 | ~ spl30_51 | spl30_120)), inference(subsumption_resolution, [], [f4753, f648])).
fof(f4753, plain, (~ (e2 = op(e0, e3)) | (~ spl30_9 | ~ spl30_44 | spl30_120)), inference(backward_demodulation, [], [f4675, f470])).
fof(f4675, plain, (~ (e2 = op(op(e3, e1), e3)) | (~ spl30_44 | spl30_120)), inference(forward_demodulation, [], [f1160, f618])).
fof(f618, plain, ((e3 = op(e1, e1)) | ~ spl30_44), inference(avatar_component_clause, [], [f616])).
fof(f4754, plain, (~ spl30_5 | ~ spl30_9), inference(avatar_split_clause, [], [f4748, f468, f451])).
fof(f4748, plain, (~ (e0 = op(e3, e2)) | ~ spl30_9), inference(backward_demodulation, [], [f208, f470])).
fof(f4747, plain, (~ spl30_5 | ~ spl30_21), inference(avatar_split_clause, [], [f4738, f519, f451])).
fof(f4738, plain, (~ (e0 = op(e3, e2)) | ~ spl30_21), inference(backward_demodulation, [], [f180, f521])).
fof(f521, plain, ((e0 = op(e2, e2)) | ~ spl30_21), inference(avatar_component_clause, [], [f519])).
fof(f4737, plain, (~ spl30_23 | ~ spl30_27), inference(avatar_split_clause, [], [f4734, f544, f527])).
fof(f4734, plain, (~ (e2 = op(e2, e2)) | ~ spl30_27), inference(backward_demodulation, [], [f202, f546])).
fof(f202, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f4723, plain, (~ spl30_35 | ~ spl30_47), inference(avatar_split_clause, [], [f4721, f629, f578])).
fof(f578, plain, (spl30_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl30_35])])).
fof(f4721, plain, (~ (e2 = op(e1, e3)) | ~ spl30_47), inference(backward_demodulation, [], [f195, f631])).
fof(f4719, plain, (~ spl30_35 | ~ spl30_51), inference(avatar_split_clause, [], [f4716, f646, f578])).
fof(f4716, plain, (~ (e2 = op(e1, e3)) | ~ spl30_51), inference(backward_demodulation, [], [f181, f648])).
fof(f4705, plain, (~ spl30_45 | ~ spl30_61), inference(avatar_split_clause, [], [f4693, f689, f621])).
fof(f621, plain, (spl30_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_45])])).
fof(f4693, plain, (~ (e0 = op(e1, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f163, f691])).
fof(f4692, plain, (~ spl30_25 | ~ spl30_30 | spl30_69), inference(avatar_split_clause, [], [f4691, f757, f557, f536])).
fof(f4691, plain, (~ (e0 = op(e2, e1)) | (~ spl30_30 | spl30_69)), inference(forward_demodulation, [], [f759, f559])).
fof(f559, plain, ((e1 = op(e2, e0)) | ~ spl30_30), inference(avatar_component_clause, [], [f557])).
fof(f759, plain, (~ (e0 = op(e2, op(e2, e0))) | spl30_69), inference(avatar_component_clause, [], [f757])).
fof(f4687, plain, (~ spl30_11 | ~ spl30_44 | spl30_118), inference(avatar_split_clause, [], [f4686, f1148, f616, f476])).
fof(f4686, plain, (~ (e2 = op(e3, e1)) | (~ spl30_44 | spl30_118)), inference(forward_demodulation, [], [f1150, f618])).
fof(f4681, plain, (~ spl30_35 | ~ spl30_2 | spl30_131), inference(avatar_split_clause, [], [f4680, f1213, f438, f578])).
fof(f4680, plain, (~ (e2 = op(e1, e3)) | (~ spl30_2 | spl30_131)), inference(forward_demodulation, [], [f1215, f440])).
fof(f440, plain, ((e1 = op(e3, e3)) | ~ spl30_2), inference(avatar_component_clause, [], [f438])).
fof(f4659, plain, (~ spl30_12 | ~ spl30_2 | spl30_81), inference(avatar_split_clause, [], [f4641, f820, f438, f480])).
fof(f820, plain, (spl30_81 <=> (e3 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_81])])).
fof(f4641, plain, (~ (e3 = op(e3, e1)) | (~ spl30_2 | spl30_81)), inference(backward_demodulation, [], [f822, f440])).
fof(f822, plain, (~ (e3 = op(e3, op(e3, e3))) | spl30_81), inference(avatar_component_clause, [], [f820])).
fof(f4658, plain, (~ spl30_6 | ~ spl30_2), inference(avatar_split_clause, [], [f4657, f438, f455])).
fof(f4657, plain, (~ (e1 = op(e3, e2)) | ~ spl30_2), inference(forward_demodulation, [], [f210, f440])).
fof(f210, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4654, plain, (~ spl30_6 | ~ spl30_38), inference(avatar_split_clause, [], [f4653, f591, f455])).
fof(f4653, plain, (~ (e1 = op(e3, e2)) | ~ spl30_38), inference(forward_demodulation, [], [f179, f593])).
fof(f593, plain, ((e1 = op(e1, e2)) | ~ spl30_38), inference(avatar_component_clause, [], [f591])).
fof(f4652, plain, (~ spl30_8 | ~ spl30_56), inference(avatar_split_clause, [], [f4651, f667, f463])).
fof(f4651, plain, (~ (e3 = op(e3, e2)) | ~ spl30_56), inference(forward_demodulation, [], [f177, f669])).
fof(f669, plain, ((e3 = op(e0, e2)) | ~ spl30_56), inference(avatar_component_clause, [], [f667])).
fof(f4650, plain, (~ spl30_2 | ~ spl30_25 | spl30_129 | ~ spl30_131), inference(avatar_contradiction_clause, [], [f4649])).
fof(f4649, plain, ($false | (~ spl30_2 | ~ spl30_25 | spl30_129 | ~ spl30_131)), inference(subsumption_resolution, [], [f4646, f538])).
fof(f4646, plain, (~ (e0 = op(e2, e1)) | (~ spl30_2 | spl30_129 | ~ spl30_131)), inference(backward_demodulation, [], [f4507, f440])).
fof(f4507, plain, (~ (e0 = op(e2, op(e3, e3))) | (spl30_129 | ~ spl30_131)), inference(forward_demodulation, [], [f1206, f1214])).
fof(f4637, plain, (~ spl30_3 | ~ spl30_11), inference(avatar_split_clause, [], [f4634, f476, f442])).
fof(f4634, plain, (~ (e2 = op(e3, e3)) | ~ spl30_11), inference(backward_demodulation, [], [f209, f478])).
fof(f4633, plain, (~ spl30_12 | ~ spl30_16), inference(avatar_split_clause, [], [f4630, f497, f480])).
fof(f4630, plain, (~ (e3 = op(e3, e1)) | ~ spl30_16), inference(backward_demodulation, [], [f205, f499])).
fof(f205, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f4612, plain, (~ spl30_17 | ~ spl30_25), inference(avatar_split_clause, [], [f4607, f536, f502])).
fof(f4607, plain, (~ (e0 = op(e2, e3)) | ~ spl30_25), inference(backward_demodulation, [], [f203, f538])).
fof(f4604, plain, (spl30_25 | ~ spl30_30 | ~ spl30_69), inference(avatar_split_clause, [], [f4601, f757, f557, f536])).
fof(f4601, plain, ((e0 = op(e2, e1)) | (~ spl30_30 | ~ spl30_69)), inference(backward_demodulation, [], [f758, f559])).
fof(f4603, plain, (~ spl30_26 | ~ spl30_30), inference(avatar_split_clause, [], [f4596, f557, f540])).
fof(f4596, plain, (~ (e1 = op(e2, e1)) | ~ spl30_30), inference(backward_demodulation, [], [f199, f559])).
fof(f4602, plain, (~ spl30_14 | ~ spl30_30), inference(avatar_split_clause, [], [f4595, f557, f489])).
fof(f4595, plain, (~ (e1 = op(e3, e0)) | ~ spl30_30), inference(backward_demodulation, [], [f168, f559])).
fof(f4593, plain, (~ spl30_3 | ~ spl30_35), inference(avatar_split_clause, [], [f4590, f578, f442])).
fof(f4590, plain, (~ (e2 = op(e3, e3)) | ~ spl30_35), inference(backward_demodulation, [], [f185, f580])).
fof(f580, plain, ((e2 = op(e1, e3)) | ~ spl30_35), inference(avatar_component_clause, [], [f578])).
fof(f185, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f4588, plain, (~ spl30_34 | ~ spl30_38), inference(avatar_split_clause, [], [f4585, f591, f574])).
fof(f4585, plain, (~ (e1 = op(e1, e3)) | ~ spl30_38), inference(backward_demodulation, [], [f198, f593])).
fof(f4580, plain, (~ spl30_34 | ~ spl30_44 | spl30_73), inference(avatar_split_clause, [], [f4573, f778, f616, f574])).
fof(f778, plain, (spl30_73 <=> (e1 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_73])])).
fof(f4573, plain, (~ (e1 = op(e1, e3)) | (~ spl30_44 | spl30_73)), inference(backward_demodulation, [], [f780, f618])).
fof(f780, plain, (~ (e1 = op(e1, op(e1, e1))) | spl30_73), inference(avatar_component_clause, [], [f778])).
fof(f4579, plain, (~ spl30_12 | ~ spl30_44), inference(avatar_split_clause, [], [f4570, f616, f480])).
fof(f4570, plain, (~ (e3 = op(e3, e1)) | ~ spl30_44), inference(backward_demodulation, [], [f173, f618])).
fof(f4568, plain, (~ spl30_41 | ~ spl30_45), inference(avatar_split_clause, [], [f4561, f621, f604])).
fof(f4561, plain, (~ (e0 = op(e1, e1)) | ~ spl30_45), inference(backward_demodulation, [], [f193, f623])).
fof(f623, plain, ((e0 = op(e1, e0)) | ~ spl30_45), inference(avatar_component_clause, [], [f621])).
fof(f4555, plain, (~ spl30_17 | ~ spl30_49), inference(avatar_split_clause, [], [f4549, f638, f502])).
fof(f4549, plain, (~ (e0 = op(e2, e3)) | ~ spl30_49), inference(backward_demodulation, [], [f182, f640])).
fof(f4547, plain, (~ spl30_52 | ~ spl30_56), inference(avatar_split_clause, [], [f4544, f667, f650])).
fof(f4544, plain, (~ (e3 = op(e0, e3)) | ~ spl30_56), inference(backward_demodulation, [], [f192, f669])).
fof(f4546, plain, (~ spl30_24 | ~ spl30_56), inference(avatar_split_clause, [], [f4543, f667, f531])).
fof(f4543, plain, (~ (e3 = op(e2, e2)) | ~ spl30_56), inference(backward_demodulation, [], [f176, f669])).
fof(f4541, plain, (~ spl30_54 | ~ spl30_58), inference(avatar_split_clause, [], [f4535, f676, f659])).
fof(f4535, plain, (~ (e1 = op(e0, e2)) | ~ spl30_58), inference(backward_demodulation, [], [f190, f678])).
fof(f4540, plain, (~ spl30_26 | ~ spl30_58), inference(avatar_split_clause, [], [f4533, f676, f540])).
fof(f4533, plain, (~ (e1 = op(e2, e1)) | ~ spl30_58), inference(backward_demodulation, [], [f170, f678])).
fof(f4529, plain, (~ spl30_59 | ~ spl30_63), inference(avatar_split_clause, [], [f4521, f697, f680])).
fof(f4521, plain, (~ (e2 = op(e0, e1)) | ~ spl30_63), inference(backward_demodulation, [], [f187, f699])).
fof(f4528, plain, (~ spl30_31 | ~ spl30_63), inference(avatar_split_clause, [], [f4519, f697, f561])).
fof(f4519, plain, (~ (e2 = op(e2, e0)) | ~ spl30_63), inference(backward_demodulation, [], [f164, f699])).
fof(f4504, plain, (~ spl30_3 | ~ spl30_54 | spl30_132 | ~ spl30_133), inference(avatar_contradiction_clause, [], [f4503])).
fof(f4503, plain, ($false | (~ spl30_3 | ~ spl30_54 | spl30_132 | ~ spl30_133)), inference(subsumption_resolution, [], [f4500, f661])).
fof(f4500, plain, (~ (e1 = op(e0, e2)) | (~ spl30_3 | spl30_132 | ~ spl30_133)), inference(backward_demodulation, [], [f4400, f444])).
fof(f4400, plain, (~ (e1 = op(e0, op(e3, e3))) | (spl30_132 | ~ spl30_133)), inference(backward_demodulation, [], [f1220, f1223])).
fof(f1220, plain, (~ (e1 = op(op(op(e3, e3), e3), op(e3, e3))) | spl30_132), inference(avatar_component_clause, [], [f1218])).
fof(f1218, plain, (spl30_132 <=> (e1 = op(op(op(e3, e3), e3), op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_132])])).
fof(f4495, plain, (~ spl30_4 | ~ spl30_12), inference(avatar_split_clause, [], [f4493, f480, f446])).
fof(f4493, plain, (~ (e3 = op(e3, e3)) | ~ spl30_12), inference(backward_demodulation, [], [f209, f482])).
fof(f4492, plain, (~ spl30_10 | ~ spl30_14), inference(avatar_split_clause, [], [f4487, f489, f472])).
fof(f4487, plain, (~ (e1 = op(e3, e1)) | ~ spl30_14), inference(backward_demodulation, [], [f205, f491])).
fof(f4473, plain, (~ spl30_27 | ~ spl30_31), inference(avatar_split_clause, [], [f4468, f561, f544])).
fof(f4468, plain, (~ (e2 = op(e2, e1)) | ~ spl30_31), inference(backward_demodulation, [], [f199, f563])).
fof(f563, plain, ((e2 = op(e2, e0)) | ~ spl30_31), inference(avatar_component_clause, [], [f561])).
fof(f4472, plain, (~ spl30_15 | ~ spl30_31), inference(avatar_split_clause, [], [f4467, f561, f493])).
fof(f4467, plain, (~ (e2 = op(e3, e0)) | ~ spl30_31), inference(backward_demodulation, [], [f168, f563])).
fof(f4459, plain, (~ spl30_35 | ~ spl30_39), inference(avatar_split_clause, [], [f4456, f595, f578])).
fof(f4456, plain, (~ (e2 = op(e1, e3)) | ~ spl30_39), inference(backward_demodulation, [], [f198, f597])).
fof(f597, plain, ((e2 = op(e1, e2)) | ~ spl30_39), inference(avatar_component_clause, [], [f595])).
fof(f4455, plain, (~ spl30_4 | ~ spl30_52), inference(avatar_split_clause, [], [f4453, f650, f446])).
fof(f4453, plain, (~ (e3 = op(e3, e3)) | ~ spl30_52), inference(backward_demodulation, [], [f183, f652])).
fof(f4450, plain, (~ spl30_38 | ~ spl30_54), inference(avatar_split_clause, [], [f4445, f659, f591])).
fof(f4445, plain, (~ (e1 = op(e1, e2)) | ~ spl30_54), inference(backward_demodulation, [], [f175, f661])).
fof(f4441, plain, (~ spl30_55 | ~ spl30_59), inference(avatar_split_clause, [], [f4433, f680, f663])).
fof(f4433, plain, (~ (e2 = op(e0, e2)) | ~ spl30_59), inference(backward_demodulation, [], [f190, f682])).
fof(f4440, plain, (~ spl30_27 | ~ spl30_59), inference(avatar_split_clause, [], [f4431, f680, f544])).
fof(f4431, plain, (~ (e2 = op(e2, e1)) | ~ spl30_59), inference(backward_demodulation, [], [f170, f682])).
fof(f4407, plain, (~ spl30_15 | ~ spl30_5 | spl30_82), inference(avatar_split_clause, [], [f4406, f825, f451, f493])).
fof(f825, plain, (spl30_82 <=> (e2 = op(e3, op(e3, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_82])])).
fof(f4406, plain, (~ (e2 = op(e3, e0)) | (~ spl30_5 | spl30_82)), inference(forward_demodulation, [], [f827, f453])).
fof(f453, plain, ((e0 = op(e3, e2)) | ~ spl30_5), inference(avatar_component_clause, [], [f451])).
fof(f827, plain, (~ (e2 = op(e3, op(e3, e2))) | spl30_82), inference(avatar_component_clause, [], [f825])).
fof(f4392, plain, (~ spl30_50 | ~ spl30_5 | ~ spl30_24 | spl30_126), inference(avatar_split_clause, [], [f4383, f1188, f531, f451, f642])).
fof(f1188, plain, (spl30_126 <=> (e1 = op(op(op(e2, e2), e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_126])])).
fof(f4383, plain, (~ (e1 = op(e0, e3)) | (~ spl30_5 | ~ spl30_24 | spl30_126)), inference(backward_demodulation, [], [f4311, f453])).
fof(f4311, plain, (~ (e1 = op(op(e3, e2), e3)) | (~ spl30_24 | spl30_126)), inference(forward_demodulation, [], [f1190, f533])).
fof(f1190, plain, (~ (e1 = op(op(op(e2, e2), e2), op(e2, e2))) | spl30_126), inference(avatar_component_clause, [], [f1188])).
fof(f4370, plain, (~ spl30_11 | ~ spl30_15), inference(avatar_split_clause, [], [f4366, f493, f476])).
fof(f4366, plain, (~ (e2 = op(e3, e1)) | ~ spl30_15), inference(backward_demodulation, [], [f205, f495])).
fof(f4319, plain, (~ spl30_11 | ~ spl30_48 | ~ spl30_62 | spl30_114), inference(avatar_split_clause, [], [f4318, f1128, f693, f633, f476])).
fof(f1128, plain, (spl30_114 <=> (e2 = op(op(op(e0, e0), e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_114])])).
fof(f4318, plain, (~ (e2 = op(e3, e1)) | (~ spl30_48 | ~ spl30_62 | spl30_114)), inference(forward_demodulation, [], [f4317, f635])).
fof(f635, plain, ((e3 = op(e1, e0)) | ~ spl30_48), inference(avatar_component_clause, [], [f633])).
fof(f4317, plain, (~ (e2 = op(op(e1, e0), e1)) | (~ spl30_62 | spl30_114)), inference(forward_demodulation, [], [f1130, f695])).
fof(f1130, plain, (~ (e2 = op(op(op(e0, e0), e0), op(e0, e0))) | spl30_114), inference(avatar_component_clause, [], [f1128])).
fof(f4313, plain, (~ spl30_6 | ~ spl30_24 | spl30_124), inference(avatar_split_clause, [], [f4312, f1178, f531, f455])).
fof(f1178, plain, (spl30_124 <=> (e1 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_124])])).
fof(f4312, plain, (~ (e1 = op(e3, e2)) | (~ spl30_24 | spl30_124)), inference(forward_demodulation, [], [f1180, f533])).
fof(f1180, plain, (~ (e1 = op(op(e2, e2), e2)) | spl30_124), inference(avatar_component_clause, [], [f1178])).
fof(f4282, plain, (~ spl30_32 | ~ spl30_48), inference(avatar_split_clause, [], [f4281, f633, f565])).
fof(f4281, plain, (~ (e3 = op(e2, e0)) | ~ spl30_48), inference(forward_demodulation, [], [f166, f635])).
fof(f4270, plain, (spl30_2 | ~ spl30_4 | ~ spl30_130), inference(avatar_contradiction_clause, [], [f4269])).
fof(f4269, plain, ($false | (spl30_2 | ~ spl30_4 | ~ spl30_130)), inference(subsumption_resolution, [], [f4266, f439])).
fof(f439, plain, (~ (e1 = op(e3, e3)) | spl30_2), inference(avatar_component_clause, [], [f438])).
fof(f4266, plain, ((e1 = op(e3, e3)) | (~ spl30_4 | ~ spl30_130)), inference(backward_demodulation, [], [f1209, f448])).
fof(f4238, plain, (~ spl30_20 | ~ spl30_24), inference(avatar_split_clause, [], [f4232, f531, f514])).
fof(f4232, plain, (~ (e3 = op(e2, e3)) | ~ spl30_24), inference(backward_demodulation, [], [f204, f533])).
fof(f4230, plain, (~ spl30_22 | ~ spl30_26), inference(avatar_split_clause, [], [f4225, f540, f523])).
fof(f4225, plain, (~ (e1 = op(e2, e2)) | ~ spl30_26), inference(backward_demodulation, [], [f202, f542])).
fof(f4216, plain, (~ spl30_7 | ~ spl30_39), inference(avatar_split_clause, [], [f4212, f595, f459])).
fof(f4212, plain, (~ (e2 = op(e3, e2)) | ~ spl30_39), inference(backward_demodulation, [], [f179, f597])).
fof(f4209, plain, (~ spl30_25 | ~ spl30_41), inference(avatar_split_clause, [], [f4198, f604, f536])).
fof(f4198, plain, (~ (e0 = op(e2, e1)) | ~ spl30_41), inference(backward_demodulation, [], [f172, f606])).
fof(f4196, plain, (~ spl30_40 | ~ spl30_48), inference(avatar_split_clause, [], [f4191, f633, f599])).
fof(f4191, plain, (~ (e3 = op(e1, e2)) | ~ spl30_48), inference(backward_demodulation, [], [f194, f635])).
fof(f4180, plain, (~ spl30_12 | ~ spl30_60), inference(avatar_split_clause, [], [f4177, f684, f480])).
fof(f4177, plain, (~ (e3 = op(e3, e1)) | ~ spl30_60), inference(backward_demodulation, [], [f171, f686])).
fof(f4174, plain, (~ spl30_48 | ~ spl30_62 | spl30_113), inference(avatar_split_clause, [], [f4170, f1123, f693, f633])).
fof(f1123, plain, (spl30_113 <=> (e3 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_113])])).
fof(f4170, plain, (~ (e3 = op(e1, e0)) | (~ spl30_62 | spl30_113)), inference(backward_demodulation, [], [f1125, f695])).
fof(f1125, plain, (~ (e3 = op(op(e0, e0), e0)) | spl30_113), inference(avatar_component_clause, [], [f1123])).
fof(f4104, plain, (~ spl30_21 | ~ spl30_25), inference(avatar_split_clause, [], [f4098, f536, f519])).
fof(f4098, plain, (~ (e0 = op(e2, e2)) | ~ spl30_25), inference(backward_demodulation, [], [f202, f538])).
fof(f4096, plain, (~ spl30_21 | ~ spl30_31 | spl30_69), inference(avatar_split_clause, [], [f4094, f757, f561, f519])).
fof(f4094, plain, (~ (e0 = op(e2, e2)) | (~ spl30_31 | spl30_69)), inference(backward_demodulation, [], [f759, f563])).
fof(f4095, plain, (~ spl30_19 | ~ spl30_31), inference(avatar_split_clause, [], [f4092, f561, f510])).
fof(f4092, plain, (~ (e2 = op(e2, e3)) | ~ spl30_31), inference(backward_demodulation, [], [f201, f563])).
fof(f4088, plain, (~ spl30_19 | ~ spl30_35), inference(avatar_split_clause, [], [f4084, f578, f510])).
fof(f4084, plain, (~ (e2 = op(e2, e3)) | ~ spl30_35), inference(backward_demodulation, [], [f184, f580])).
fof(f4077, plain, (~ spl30_33 | ~ spl30_45), inference(avatar_split_clause, [], [f4071, f621, f570])).
fof(f4071, plain, (~ (e0 = op(e1, e3)) | ~ spl30_45), inference(backward_demodulation, [], [f195, f623])).
fof(f4076, plain, (~ spl30_13 | ~ spl30_45), inference(avatar_split_clause, [], [f4069, f621, f485])).
fof(f4069, plain, (~ (e0 = op(e3, e0)) | ~ spl30_45), inference(backward_demodulation, [], [f167, f623])).
fof(f4059, plain, (~ spl30_21 | ~ spl30_53), inference(avatar_split_clause, [], [f4052, f655, f519])).
fof(f4052, plain, (~ (e0 = op(e2, e2)) | ~ spl30_53), inference(backward_demodulation, [], [f176, f657])).
fof(f4049, plain, (~ spl30_11 | ~ spl30_59), inference(avatar_split_clause, [], [f4044, f680, f476])).
fof(f4044, plain, (~ (e2 = op(e3, e1)) | ~ spl30_59), inference(backward_demodulation, [], [f171, f682])).
fof(f4042, plain, (~ spl30_35 | ~ spl30_64 | spl30_114 | ~ spl30_115), inference(avatar_split_clause, [], [f4038, f1132, f1128, f701, f578])).
fof(f1132, plain, (spl30_115 <=> (e1 = op(op(e0, e0), e0))), introduced(avatar_definition, [new_symbols(naming, [spl30_115])])).
fof(f4038, plain, (~ (e2 = op(e1, e3)) | (~ spl30_64 | spl30_114 | ~ spl30_115)), inference(backward_demodulation, [], [f4027, f703])).
fof(f4027, plain, (~ (e2 = op(e1, op(e0, e0))) | (spl30_114 | ~ spl30_115)), inference(forward_demodulation, [], [f1130, f1133])).
fof(f1133, plain, ((e1 = op(op(e0, e0), e0)) | ~ spl30_115), inference(avatar_component_clause, [], [f1132])).
fof(f4039, plain, (~ spl30_48 | ~ spl30_64), inference(avatar_split_clause, [], [f4028, f701, f633])).
fof(f4028, plain, (~ (e3 = op(e1, e0)) | ~ spl30_64), inference(backward_demodulation, [], [f163, f703])).
fof(f4004, plain, (spl30_33 | ~ spl30_48 | ~ spl30_74), inference(avatar_split_clause, [], [f4001, f783, f633, f570])).
fof(f4001, plain, ((e0 = op(e1, e3)) | (~ spl30_48 | ~ spl30_74)), inference(backward_demodulation, [], [f784, f635])).
fof(f4003, plain, (~ spl30_36 | ~ spl30_48), inference(avatar_split_clause, [], [f4000, f633, f582])).
fof(f4000, plain, (~ (e3 = op(e1, e3)) | ~ spl30_48), inference(backward_demodulation, [], [f195, f635])).
fof(f3995, plain, (~ spl30_49 | ~ spl30_57), inference(avatar_split_clause, [], [f3990, f672, f638])).
fof(f3990, plain, (~ (e0 = op(e0, e3)) | ~ spl30_57), inference(backward_demodulation, [], [f191, f674])).
fof(f3989, plain, (~ spl30_21 | ~ spl30_54 | spl30_124), inference(avatar_contradiction_clause, [], [f3988])).
fof(f3988, plain, ($false | (~ spl30_21 | ~ spl30_54 | spl30_124)), inference(subsumption_resolution, [], [f3987, f661])).
fof(f3987, plain, (~ (e1 = op(e0, e2)) | (~ spl30_21 | spl30_124)), inference(forward_demodulation, [], [f1180, f521])).
fof(f3944, plain, (~ spl30_9 | ~ spl30_11), inference(avatar_contradiction_clause, [], [f3943])).
fof(f3943, plain, ($false | (~ spl30_9 | ~ spl30_11)), inference(subsumption_resolution, [], [f3942, f212])).
fof(f3942, plain, ((e0 = e2) | (~ spl30_9 | ~ spl30_11)), inference(backward_demodulation, [], [f478, f470])).
fof(f3923, plain, (~ spl30_21 | spl30_53 | ~ spl30_127), inference(avatar_contradiction_clause, [], [f3922])).
fof(f3922, plain, ($false | (~ spl30_21 | spl30_53 | ~ spl30_127)), inference(subsumption_resolution, [], [f3918, f656])).
fof(f656, plain, (~ (e0 = op(e0, e2)) | spl30_53), inference(avatar_component_clause, [], [f655])).
fof(f3918, plain, ((e0 = op(e0, e2)) | (~ spl30_21 | ~ spl30_127)), inference(backward_demodulation, [], [f1193, f521])).
fof(f3904, plain, (~ spl30_22 | ~ spl30_30), inference(avatar_split_clause, [], [f3898, f557, f523])).
fof(f3898, plain, (~ (e1 = op(e2, e2)) | ~ spl30_30), inference(backward_demodulation, [], [f200, f559])).
fof(f3880, plain, (~ spl30_34 | ~ spl30_42), inference(avatar_split_clause, [], [f3872, f608, f574])).
fof(f3872, plain, (~ (e1 = op(e1, e3)) | ~ spl30_42), inference(backward_demodulation, [], [f197, f610])).
fof(f197, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3868, plain, (~ spl30_37 | ~ spl30_45), inference(avatar_split_clause, [], [f3862, f621, f587])).
fof(f3862, plain, (~ (e0 = op(e1, e2)) | ~ spl30_45), inference(backward_demodulation, [], [f194, f623])).
fof(f3830, plain, (~ spl30_51 | ~ spl30_63), inference(avatar_split_clause, [], [f3823, f697, f646])).
fof(f3823, plain, (~ (e2 = op(e0, e3)) | ~ spl30_63), inference(backward_demodulation, [], [f189, f699])).
fof(f3803, plain, (~ spl30_25 | ~ spl30_43 | ~ spl30_56 | spl30_122), inference(avatar_contradiction_clause, [], [f3802])).
fof(f3802, plain, ($false | (~ spl30_25 | ~ spl30_43 | ~ spl30_56 | spl30_122)), inference(subsumption_resolution, [], [f3801, f669])).
fof(f3801, plain, (~ (e3 = op(e0, e2)) | (~ spl30_25 | ~ spl30_43 | spl30_122)), inference(forward_demodulation, [], [f3628, f538])).
fof(f3628, plain, (~ (e3 = op(op(e2, e1), e2)) | (~ spl30_43 | spl30_122)), inference(forward_demodulation, [], [f1170, f614])).
fof(f3788, plain, (~ spl30_4 | ~ spl30_20), inference(avatar_split_clause, [], [f3787, f514, f446])).
fof(f3787, plain, (~ (e3 = op(e3, e3)) | ~ spl30_20), inference(backward_demodulation, [], [f186, f516])).
fof(f516, plain, ((e3 = op(e2, e3)) | ~ spl30_20), inference(avatar_component_clause, [], [f514])).
fof(f3764, plain, (~ spl30_40 | ~ spl30_56), inference(avatar_split_clause, [], [f3762, f667, f599])).
fof(f3762, plain, (~ (e3 = op(e1, e2)) | ~ spl30_56), inference(backward_demodulation, [], [f175, f669])).
fof(f3742, plain, (~ spl30_28 | ~ spl30_43 | spl30_119), inference(avatar_split_clause, [], [f3741, f1153, f612, f548])).
fof(f3741, plain, (~ (e3 = op(e2, e1)) | (~ spl30_43 | spl30_119)), inference(forward_demodulation, [], [f1155, f614])).
fof(f1155, plain, (~ (e3 = op(op(e1, e1), e1)) | spl30_119), inference(avatar_component_clause, [], [f1153])).
fof(f3738, plain, (~ spl30_40 | ~ spl30_22 | spl30_125), inference(avatar_split_clause, [], [f3737, f1183, f523, f599])).
fof(f1183, plain, (spl30_125 <=> (e3 = op(op(e2, e2), e2))), introduced(avatar_definition, [new_symbols(naming, [spl30_125])])).
fof(f3737, plain, (~ (e3 = op(e1, e2)) | (~ spl30_22 | spl30_125)), inference(forward_demodulation, [], [f1185, f525])).
fof(f1185, plain, (~ (e3 = op(op(e2, e2), e2)) | spl30_125), inference(avatar_component_clause, [], [f1183])).
fof(f3713, plain, (~ spl30_11 | ~ spl30_43), inference(avatar_split_clause, [], [f3712, f612, f476])).
fof(f3712, plain, (~ (e2 = op(e3, e1)) | ~ spl30_43), inference(forward_demodulation, [], [f173, f614])).
fof(f3711, plain, (~ spl30_10 | ~ spl30_58), inference(avatar_split_clause, [], [f3710, f676, f472])).
fof(f3710, plain, (~ (e1 = op(e3, e1)) | ~ spl30_58), inference(forward_demodulation, [], [f171, f678])).
fof(f3683, plain, (~ spl30_9 | ~ spl30_22 | ~ spl30_40 | spl30_123), inference(avatar_contradiction_clause, [], [f3682])).
fof(f3682, plain, ($false | (~ spl30_9 | ~ spl30_22 | ~ spl30_40 | spl30_123)), inference(subsumption_resolution, [], [f3678, f470])).
fof(f3678, plain, (~ (e0 = op(e3, e1)) | (~ spl30_22 | ~ spl30_40 | spl30_123)), inference(backward_demodulation, [], [f3627, f601])).
fof(f3627, plain, (~ (e0 = op(op(e1, e2), e1)) | (~ spl30_22 | spl30_123)), inference(forward_demodulation, [], [f1176, f525])).
fof(f1176, plain, (~ (e0 = op(op(op(e2, e2), e2), op(e2, e2))) | spl30_123), inference(avatar_component_clause, [], [f1174])).
fof(f1174, plain, (spl30_123 <=> (e0 = op(op(op(e2, e2), e2), op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_123])])).
fof(f3650, plain, (~ spl30_14 | ~ spl30_64 | spl30_115), inference(avatar_split_clause, [], [f3646, f1132, f701, f489])).
fof(f3646, plain, (~ (e1 = op(e3, e0)) | (~ spl30_64 | spl30_115)), inference(backward_demodulation, [], [f1134, f703])).
fof(f1134, plain, (~ (e1 = op(op(e0, e0), e0)) | spl30_115), inference(avatar_component_clause, [], [f1132])).
fof(f3649, plain, (~ spl30_15 | ~ spl30_64 | spl30_112), inference(avatar_split_clause, [], [f3644, f1118, f701, f493])).
fof(f3644, plain, (~ (e2 = op(e3, e0)) | (~ spl30_64 | spl30_112)), inference(backward_demodulation, [], [f1120, f703])).
fof(f1120, plain, (~ (e2 = op(op(e0, e0), e0)) | spl30_112), inference(avatar_component_clause, [], [f1118])).
fof(f3635, plain, (spl30_14 | ~ spl30_9 | ~ spl30_83), inference(avatar_split_clause, [], [f3634, f830, f468, f489])).
fof(f830, plain, (spl30_83 <=> (e1 = op(e3, op(e3, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl30_83])])).
fof(f3634, plain, ((e1 = op(e3, e0)) | (~ spl30_9 | ~ spl30_83)), inference(forward_demodulation, [], [f831, f470])).
fof(f831, plain, ((e1 = op(e3, op(e3, e1))) | ~ spl30_83), inference(avatar_component_clause, [], [f830])).
fof(f3615, plain, (~ spl30_37 | ~ spl30_22 | spl30_127), inference(avatar_split_clause, [], [f3588, f1192, f523, f587])).
fof(f3588, plain, (~ (e0 = op(e1, e2)) | (~ spl30_22 | spl30_127)), inference(backward_demodulation, [], [f1194, f525])).
fof(f1194, plain, (~ (e0 = op(op(e2, e2), e2)) | spl30_127), inference(avatar_component_clause, [], [f1192])).
fof(f3607, plain, (~ spl30_2 | ~ spl30_34), inference(avatar_split_clause, [], [f3606, f574, f438])).
fof(f3606, plain, (~ (e1 = op(e3, e3)) | ~ spl30_34), inference(forward_demodulation, [], [f185, f576])).
fof(f3603, plain, (~ spl30_7 | ~ spl30_15), inference(avatar_split_clause, [], [f3602, f493, f459])).
fof(f3602, plain, (~ (e2 = op(e3, e2)) | ~ spl30_15), inference(backward_demodulation, [], [f206, f495])).
fof(f3599, plain, (~ spl30_17 | ~ spl30_19), inference(avatar_contradiction_clause, [], [f3598])).
fof(f3598, plain, ($false | (~ spl30_17 | ~ spl30_19)), inference(subsumption_resolution, [], [f3597, f212])).
fof(f3597, plain, ((e0 = e2) | (~ spl30_17 | ~ spl30_19)), inference(backward_demodulation, [], [f512, f504])).
fof(f3559, plain, (~ spl30_35 | ~ spl30_43), inference(avatar_split_clause, [], [f3550, f612, f578])).
fof(f3550, plain, (~ (e2 = op(e1, e3)) | ~ spl30_43), inference(backward_demodulation, [], [f197, f614])).
fof(f3540, plain, (~ spl30_20 | ~ spl30_52), inference(avatar_split_clause, [], [f3538, f650, f514])).
fof(f3538, plain, (~ (e3 = op(e2, e3)) | ~ spl30_52), inference(backward_demodulation, [], [f182, f652])).
fof(f3528, plain, (~ spl30_42 | ~ spl30_58), inference(avatar_split_clause, [], [f3521, f676, f608])).
fof(f3521, plain, (~ (e1 = op(e1, e1)) | ~ spl30_58), inference(backward_demodulation, [], [f169, f678])).
fof(f3491, plain, (spl30_35 | ~ spl30_2 | ~ spl30_131), inference(avatar_split_clause, [], [f3415, f1213, f438, f578])).
fof(f3415, plain, ((e2 = op(e1, e3)) | (~ spl30_2 | ~ spl30_131)), inference(backward_demodulation, [], [f1214, f440])).
fof(f3478, plain, (spl30_1 | ~ spl30_16 | ~ spl30_84), inference(avatar_contradiction_clause, [], [f3477])).
fof(f3477, plain, ($false | (spl30_1 | ~ spl30_16 | ~ spl30_84)), inference(subsumption_resolution, [], [f3476, f435])).
fof(f435, plain, (~ (e0 = op(e3, e3)) | spl30_1), inference(avatar_component_clause, [], [f434])).
fof(f3476, plain, ((e0 = op(e3, e3)) | (~ spl30_16 | ~ spl30_84)), inference(backward_demodulation, [], [f836, f499])).
fof(f3474, plain, (~ spl30_35 | ~ spl30_40 | spl30_72), inference(avatar_contradiction_clause, [], [f3473])).
fof(f3473, plain, ($false | (~ spl30_35 | ~ spl30_40 | spl30_72)), inference(subsumption_resolution, [], [f3472, f580])).
fof(f3472, plain, (~ (e2 = op(e1, e3)) | (~ spl30_40 | spl30_72)), inference(backward_demodulation, [], [f775, f601])).
fof(f775, plain, (~ (e2 = op(e1, op(e1, e2))) | spl30_72), inference(avatar_component_clause, [], [f773])).
fof(f773, plain, (spl30_72 <=> (e2 = op(e1, op(e1, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_72])])).
fof(f3465, plain, (~ spl30_10 | ~ spl30_42), inference(avatar_split_clause, [], [f3457, f608, f472])).
fof(f3457, plain, (~ (e1 = op(e3, e1)) | ~ spl30_42), inference(backward_demodulation, [], [f173, f610])).
fof(f3443, plain, (~ spl30_48 | ~ spl30_21 | ~ spl30_54 | spl30_128), inference(avatar_split_clause, [], [f3392, f1198, f659, f519, f633])).
fof(f3392, plain, (~ (e3 = op(e1, e0)) | (~ spl30_21 | ~ spl30_54 | spl30_128)), inference(forward_demodulation, [], [f3391, f661])).
fof(f3391, plain, (~ (e3 = op(op(e0, e2), e0)) | (~ spl30_21 | spl30_128)), inference(forward_demodulation, [], [f1200, f521])).
fof(f3431, plain, (~ spl30_40 | ~ spl30_30 | ~ spl30_63 | spl30_116), inference(avatar_split_clause, [], [f3400, f1138, f697, f557, f599])).
fof(f1138, plain, (spl30_116 <=> (e3 = op(op(op(e0, e0), e0), op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl30_116])])).
fof(f3400, plain, (~ (e3 = op(e1, e2)) | (~ spl30_30 | ~ spl30_63 | spl30_116)), inference(forward_demodulation, [], [f3399, f559])).
fof(f3399, plain, (~ (e3 = op(op(e2, e0), e2)) | (~ spl30_63 | spl30_116)), inference(forward_demodulation, [], [f1140, f699])).
fof(f1140, plain, (~ (e3 = op(op(op(e0, e0), e0), op(e0, e0))) | spl30_116), inference(avatar_component_clause, [], [f1138])).
fof(f3425, plain, (~ spl30_10 | ~ spl30_2), inference(avatar_split_clause, [], [f3424, f438, f472])).
fof(f3424, plain, (~ (e1 = op(e3, e1)) | ~ spl30_2), inference(forward_demodulation, [], [f209, f440])).
fof(f3423, plain, (spl30_12 | ~ spl30_2 | ~ spl30_81), inference(avatar_split_clause, [], [f3414, f820, f438, f480])).
fof(f3414, plain, ((e3 = op(e3, e1)) | (~ spl30_2 | ~ spl30_81)), inference(backward_demodulation, [], [f821, f440])).
fof(f821, plain, ((e3 = op(e3, op(e3, e3))) | ~ spl30_81), inference(avatar_component_clause, [], [f820])).
fof(f3411, plain, (~ spl30_7 | ~ spl30_8), inference(avatar_contradiction_clause, [], [f3410])).
fof(f3410, plain, ($false | (~ spl30_7 | ~ spl30_8)), inference(subsumption_resolution, [], [f3409, f216])).
fof(f216, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f3409, plain, ((e2 = e3) | (~ spl30_7 | ~ spl30_8)), inference(forward_demodulation, [], [f465, f461])).
fof(f461, plain, ((e2 = op(e3, e2)) | ~ spl30_7), inference(avatar_component_clause, [], [f459])).
fof(f465, plain, ((e3 = op(e3, e2)) | ~ spl30_8), inference(avatar_component_clause, [], [f463])).
fof(f3359, plain, (~ spl30_5 | ~ spl30_13), inference(avatar_split_clause, [], [f3354, f485, f451])).
fof(f3354, plain, (~ (e0 = op(e3, e2)) | ~ spl30_13), inference(backward_demodulation, [], [f206, f487])).
fof(f3325, plain, (~ spl30_18 | ~ spl30_30), inference(avatar_split_clause, [], [f3319, f557, f506])).
fof(f3319, plain, (~ (e1 = op(e2, e3)) | ~ spl30_30), inference(backward_demodulation, [], [f201, f559])).
fof(f3315, plain, (~ spl30_34 | ~ spl30_35), inference(avatar_contradiction_clause, [], [f3314])).
fof(f3314, plain, ($false | (~ spl30_34 | ~ spl30_35)), inference(subsumption_resolution, [], [f3313, f214])).
fof(f214, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f3313, plain, ((e1 = e2) | (~ spl30_34 | ~ spl30_35)), inference(backward_demodulation, [], [f580, f576])).
fof(f3301, plain, (~ spl30_44 | ~ spl30_48), inference(avatar_split_clause, [], [f3299, f633, f616])).
fof(f3299, plain, (~ (e3 = op(e1, e1)) | ~ spl30_48), inference(backward_demodulation, [], [f193, f635])).
fof(f3296, plain, (~ spl30_54 | ~ spl30_56), inference(avatar_contradiction_clause, [], [f3295])).
fof(f3295, plain, ($false | (~ spl30_54 | ~ spl30_56)), inference(subsumption_resolution, [], [f3294, f215])).
fof(f3294, plain, ((e1 = e3) | (~ spl30_54 | ~ spl30_56)), inference(backward_demodulation, [], [f669, f661])).
fof(f3264, plain, (~ spl30_30 | ~ spl30_63 | spl30_115), inference(avatar_split_clause, [], [f3119, f1132, f697, f557])).
fof(f3119, plain, (~ (e1 = op(e2, e0)) | (~ spl30_63 | spl30_115)), inference(backward_demodulation, [], [f1134, f699])).
fof(f3245, plain, (~ spl30_21 | ~ spl30_23 | spl30_127), inference(avatar_split_clause, [], [f3236, f1192, f527, f519])).
fof(f3236, plain, (~ (e0 = op(e2, e2)) | (~ spl30_23 | spl30_127)), inference(backward_demodulation, [], [f1194, f529])).
fof(f3241, plain, (~ spl30_23 | spl30_67), inference(avatar_contradiction_clause, [], [f3240])).
fof(f3240, plain, ($false | (~ spl30_23 | spl30_67)), inference(subsumption_resolution, [], [f3234, f529])).
fof(f3234, plain, (~ (e2 = op(e2, e2)) | (~ spl30_23 | spl30_67)), inference(backward_demodulation, [], [f749, f529])).
fof(f3239, plain, (~ spl30_7 | ~ spl30_23), inference(avatar_split_clause, [], [f3232, f527, f459])).
fof(f3232, plain, (~ (e2 = op(e3, e2)) | ~ spl30_23), inference(backward_demodulation, [], [f180, f529])).
fof(f3220, plain, (~ spl30_9 | ~ spl30_44 | spl30_121), inference(avatar_split_clause, [], [f3219, f1162, f616, f468])).
fof(f3219, plain, (~ (e0 = op(e3, e1)) | (~ spl30_44 | spl30_121)), inference(forward_demodulation, [], [f1164, f618])).
fof(f3213, plain, (~ spl30_22 | ~ spl30_38), inference(avatar_split_clause, [], [f3212, f591, f523])).
fof(f3212, plain, (~ (e1 = op(e2, e2)) | ~ spl30_38), inference(forward_demodulation, [], [f178, f593])).
fof(f3196, plain, (~ spl30_14 | ~ spl30_21 | ~ spl30_56 | spl30_126), inference(avatar_contradiction_clause, [], [f3195])).
fof(f3195, plain, ($false | (~ spl30_14 | ~ spl30_21 | ~ spl30_56 | spl30_126)), inference(subsumption_resolution, [], [f3194, f491])).
fof(f3194, plain, (~ (e1 = op(e3, e0)) | (~ spl30_21 | ~ spl30_56 | spl30_126)), inference(forward_demodulation, [], [f3193, f669])).
fof(f3193, plain, (~ (e1 = op(op(e0, e2), e0)) | (~ spl30_21 | spl30_126)), inference(forward_demodulation, [], [f1190, f521])).
fof(f3176, plain, (~ spl30_2 | ~ spl30_18), inference(avatar_split_clause, [], [f3172, f506, f438])).
fof(f3172, plain, (~ (e1 = op(e3, e3)) | ~ spl30_18), inference(backward_demodulation, [], [f186, f508])).
fof(f508, plain, ((e1 = op(e2, e3)) | ~ spl30_18), inference(avatar_component_clause, [], [f506])).
fof(f3148, plain, (~ spl30_33 | ~ spl30_49), inference(avatar_split_clause, [], [f3142, f638, f570])).
fof(f3142, plain, (~ (e0 = op(e1, e3)) | ~ spl30_49), inference(backward_demodulation, [], [f181, f640])).
fof(f3122, plain, (~ spl30_15 | ~ spl30_63), inference(avatar_split_clause, [], [f3111, f697, f493])).
fof(f3111, plain, (~ (e2 = op(e3, e0)) | ~ spl30_63), inference(backward_demodulation, [], [f165, f699])).
fof(f3103, plain, (~ spl30_56 | ~ spl30_21 | spl30_125), inference(avatar_split_clause, [], [f3085, f1183, f519, f667])).
fof(f3085, plain, (~ (e3 = op(e0, e2)) | (~ spl30_21 | spl30_125)), inference(forward_demodulation, [], [f1185, f521])).
fof(f3093, plain, (~ spl30_14 | ~ spl30_9 | spl30_83), inference(avatar_split_clause, [], [f3070, f830, f468, f489])).
fof(f3070, plain, (~ (e1 = op(e3, e0)) | (~ spl30_9 | spl30_83)), inference(backward_demodulation, [], [f832, f470])).
fof(f832, plain, (~ (e1 = op(e3, op(e3, e1))) | spl30_83), inference(avatar_component_clause, [], [f830])).
fof(f3082, plain, (~ spl30_2 | ~ spl30_33 | spl30_133), inference(avatar_contradiction_clause, [], [f3081])).
fof(f3081, plain, ($false | (~ spl30_2 | ~ spl30_33 | spl30_133)), inference(subsumption_resolution, [], [f3080, f572])).
fof(f3080, plain, (~ (e0 = op(e1, e3)) | (~ spl30_2 | spl30_133)), inference(forward_demodulation, [], [f1224, f440])).
fof(f1224, plain, (~ (e0 = op(op(e3, e3), e3)) | spl30_133), inference(avatar_component_clause, [], [f1222])).
fof(f3053, plain, (~ spl30_17 | ~ spl30_21), inference(avatar_split_clause, [], [f3045, f519, f502])).
fof(f3045, plain, (~ (e0 = op(e2, e3)) | ~ spl30_21), inference(backward_demodulation, [], [f204, f521])).
fof(f3044, plain, (~ spl30_2 | ~ spl30_33 | ~ spl30_59 | spl30_134), inference(avatar_contradiction_clause, [], [f3043])).
fof(f3043, plain, ($false | (~ spl30_2 | ~ spl30_33 | ~ spl30_59 | spl30_134)), inference(subsumption_resolution, [], [f3039, f682])).
fof(f3039, plain, (~ (e2 = op(e0, e1)) | (~ spl30_2 | ~ spl30_33 | spl30_134)), inference(backward_demodulation, [], [f2978, f572])).
fof(f2978, plain, (~ (e2 = op(op(e1, e3), e1)) | (~ spl30_2 | spl30_134)), inference(backward_demodulation, [], [f1230, f440])).
fof(f3032, plain, (~ spl30_23 | ~ spl30_39), inference(avatar_split_clause, [], [f3029, f595, f527])).
fof(f3029, plain, (~ (e2 = op(e2, e2)) | ~ spl30_39), inference(backward_demodulation, [], [f178, f597])).
fof(f3018, plain, (~ spl30_57 | ~ spl30_61), inference(avatar_split_clause, [], [f3008, f689, f672])).
fof(f3008, plain, (~ (e0 = op(e0, e1)) | ~ spl30_61), inference(backward_demodulation, [], [f187, f691])).
fof(f3017, plain, (~ spl30_13 | ~ spl30_61), inference(avatar_split_clause, [], [f3007, f689, f485])).
fof(f3007, plain, (~ (e0 = op(e3, e0)) | ~ spl30_61), inference(backward_demodulation, [], [f165, f691])).
fof(f3006, plain, (~ spl30_17 | ~ spl30_32 | spl30_69), inference(avatar_split_clause, [], [f3005, f757, f565, f502])).
fof(f3005, plain, (~ (e0 = op(e2, e3)) | (~ spl30_32 | spl30_69)), inference(forward_demodulation, [], [f759, f567])).
fof(f2973, plain, (~ spl30_4 | ~ spl30_8), inference(avatar_split_clause, [], [f2971, f463, f446])).
fof(f2971, plain, (~ (e3 = op(e3, e3)) | ~ spl30_8), inference(backward_demodulation, [], [f210, f465])).
fof(f2969, plain, (~ spl30_11 | ~ spl30_17 | ~ spl30_44 | spl30_117), inference(avatar_contradiction_clause, [], [f2968])).
fof(f2968, plain, ($false | (~ spl30_11 | ~ spl30_17 | ~ spl30_44 | spl30_117)), inference(subsumption_resolution, [], [f2965, f504])).
fof(f2965, plain, (~ (e0 = op(e2, e3)) | (~ spl30_11 | ~ spl30_44 | spl30_117)), inference(backward_demodulation, [], [f2812, f478])).
fof(f2812, plain, (~ (e0 = op(op(e3, e1), e3)) | (~ spl30_44 | spl30_117)), inference(forward_demodulation, [], [f1146, f618])).
fof(f2967, plain, (~ spl30_6 | ~ spl30_11 | spl30_83), inference(avatar_split_clause, [], [f2964, f830, f476, f455])).
fof(f2964, plain, (~ (e1 = op(e3, e2)) | (~ spl30_11 | spl30_83)), inference(backward_demodulation, [], [f832, f478])).
fof(f2960, plain, (~ spl30_9 | ~ spl30_13), inference(avatar_split_clause, [], [f2953, f485, f468])).
fof(f2953, plain, (~ (e0 = op(e3, e1)) | ~ spl30_13), inference(backward_demodulation, [], [f205, f487])).
fof(f2910, plain, (~ spl30_6 | ~ spl30_54), inference(avatar_split_clause, [], [f2905, f659, f455])).
fof(f2905, plain, (~ (e1 = op(e3, e2)) | ~ spl30_54), inference(backward_demodulation, [], [f177, f661])).
fof(f2894, plain, (~ spl30_6 | ~ spl30_63 | spl30_111 | ~ spl30_113), inference(avatar_split_clause, [], [f2891, f1123, f1114, f697, f455])).
fof(f2891, plain, (~ (e1 = op(e3, e2)) | (~ spl30_63 | spl30_111 | ~ spl30_113)), inference(backward_demodulation, [], [f2878, f699])).
fof(f2878, plain, (~ (e1 = op(e3, op(e0, e0))) | (spl30_111 | ~ spl30_113)), inference(forward_demodulation, [], [f1116, f1124])).
fof(f1124, plain, ((e3 = op(op(e0, e0), e0)) | ~ spl30_113), inference(avatar_component_clause, [], [f1123])).
fof(f2874, plain, (~ spl30_20 | ~ spl30_32), inference(avatar_split_clause, [], [f2873, f565, f514])).
fof(f2873, plain, (~ (e3 = op(e2, e3)) | ~ spl30_32), inference(forward_demodulation, [], [f201, f567])).
fof(f2847, plain, (~ spl30_61 | spl30_79), inference(avatar_contradiction_clause, [], [f2846])).
fof(f2846, plain, ($false | (~ spl30_61 | spl30_79)), inference(subsumption_resolution, [], [f2840, f691])).
fof(f2840, plain, (~ (e0 = op(e0, e0)) | (~ spl30_61 | spl30_79)), inference(backward_demodulation, [], [f811, f691])).
fof(f2802, plain, (~ spl30_32 | ~ spl30_63 | spl30_113), inference(avatar_contradiction_clause, [], [f2801])).
fof(f2801, plain, ($false | (~ spl30_32 | ~ spl30_63 | spl30_113)), inference(subsumption_resolution, [], [f2800, f567])).
fof(f2800, plain, (~ (e3 = op(e2, e0)) | (~ spl30_63 | spl30_113)), inference(forward_demodulation, [], [f1125, f699])).
fof(f2748, plain, (~ spl30_53 | ~ spl30_63 | spl30_79), inference(avatar_split_clause, [], [f2742, f809, f697, f655])).
fof(f2742, plain, (~ (e0 = op(e0, e2)) | (~ spl30_63 | spl30_79)), inference(backward_demodulation, [], [f811, f699])).
fof(f2736, plain, (~ spl30_24 | ~ spl30_19 | spl30_66), inference(avatar_split_clause, [], [f2735, f742, f510, f531])).
fof(f742, plain, (spl30_66 <=> (e3 = op(e2, op(e2, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_66])])).
fof(f2735, plain, (~ (e3 = op(e2, e2)) | (~ spl30_19 | spl30_66)), inference(forward_demodulation, [], [f744, f512])).
fof(f744, plain, (~ (e3 = op(e2, op(e2, e3))) | spl30_66), inference(avatar_component_clause, [], [f742])).
fof(f2708, plain, (~ spl30_60 | ~ spl30_50 | spl30_76), inference(avatar_split_clause, [], [f2617, f794, f642, f684])).
fof(f794, plain, (spl30_76 <=> (e3 = op(e0, op(e0, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_76])])).
fof(f2617, plain, (~ (e3 = op(e0, e1)) | (~ spl30_50 | spl30_76)), inference(backward_demodulation, [], [f796, f644])).
fof(f796, plain, (~ (e3 = op(e0, op(e0, e3))) | spl30_76), inference(avatar_component_clause, [], [f794])).
fof(f2686, plain, (~ spl30_2 | ~ spl30_4), inference(avatar_contradiction_clause, [], [f2685])).
fof(f2685, plain, ($false | (~ spl30_2 | ~ spl30_4)), inference(subsumption_resolution, [], [f2684, f215])).
fof(f2684, plain, ((e1 = e3) | (~ spl30_2 | ~ spl30_4)), inference(backward_demodulation, [], [f448, f440])).
fof(f2678, plain, (~ spl30_4 | spl30_81), inference(avatar_contradiction_clause, [], [f2677])).
fof(f2677, plain, ($false | (~ spl30_4 | spl30_81)), inference(subsumption_resolution, [], [f2672, f448])).
fof(f2672, plain, (~ (e3 = op(e3, e3)) | (~ spl30_4 | spl30_81)), inference(backward_demodulation, [], [f822, f448])).
fof(f2645, plain, (~ spl30_33 | ~ spl30_35), inference(avatar_contradiction_clause, [], [f2644])).
fof(f2644, plain, ($false | (~ spl30_33 | ~ spl30_35)), inference(subsumption_resolution, [], [f2643, f212])).
fof(f2643, plain, ((e0 = e2) | (~ spl30_33 | ~ spl30_35)), inference(backward_demodulation, [], [f580, f572])).
fof(f2630, plain, (~ spl30_36 | ~ spl30_44), inference(avatar_split_clause, [], [f2625, f616, f582])).
fof(f2625, plain, (~ (e3 = op(e1, e3)) | ~ spl30_44), inference(backward_demodulation, [], [f197, f618])).
fof(f2594, plain, (spl30_19 | ~ spl30_24 | ~ spl30_67), inference(avatar_split_clause, [], [f2593, f747, f531, f510])).
fof(f2593, plain, ((e2 = op(e2, e3)) | (~ spl30_24 | ~ spl30_67)), inference(forward_demodulation, [], [f748, f533])).
fof(f748, plain, ((e2 = op(e2, op(e2, e2))) | ~ spl30_67), inference(avatar_component_clause, [], [f747])).
fof(f2583, plain, (~ spl30_14 | ~ spl30_46), inference(avatar_split_clause, [], [f2582, f625, f489])).
fof(f2582, plain, (~ (e1 = op(e3, e0)) | ~ spl30_46), inference(forward_demodulation, [], [f167, f627])).
fof(f2580, plain, (~ spl30_38 | ~ spl30_46), inference(avatar_split_clause, [], [f2579, f625, f591])).
fof(f2579, plain, (~ (e1 = op(e1, e2)) | ~ spl30_46), inference(forward_demodulation, [], [f194, f627])).
fof(f2564, plain, (~ spl30_3 | ~ spl30_7), inference(avatar_split_clause, [], [f2563, f459, f442])).
fof(f2563, plain, (~ (e2 = op(e3, e3)) | ~ spl30_7), inference(forward_demodulation, [], [f210, f461])).
fof(f2556, plain, (~ spl30_6 | ~ spl30_14), inference(avatar_split_clause, [], [f2549, f489, f455])).
fof(f2549, plain, (~ (e1 = op(e3, e2)) | ~ spl30_14), inference(backward_demodulation, [], [f206, f491])).
fof(f2521, plain, (~ spl30_46 | ~ spl30_47), inference(avatar_contradiction_clause, [], [f2520])).
fof(f2520, plain, ($false | (~ spl30_46 | ~ spl30_47)), inference(subsumption_resolution, [], [f2519, f214])).
fof(f2519, plain, ((e1 = e2) | (~ spl30_46 | ~ spl30_47)), inference(backward_demodulation, [], [f631, f627])).
fof(f2510, plain, (~ spl30_30 | ~ spl30_1 | ~ spl30_51 | spl30_132), inference(avatar_split_clause, [], [f2505, f1218, f646, f434, f557])).
fof(f2505, plain, (~ (e1 = op(e2, e0)) | (~ spl30_1 | ~ spl30_51 | spl30_132)), inference(backward_demodulation, [], [f2268, f648])).
fof(f2268, plain, (~ (e1 = op(op(e0, e3), e0)) | (~ spl30_1 | spl30_132)), inference(forward_demodulation, [], [f1220, f436])).
fof(f2459, plain, (~ spl30_19 | ~ spl30_21 | ~ spl30_66), inference(avatar_contradiction_clause, [], [f2458])).
fof(f2458, plain, ($false | (~ spl30_19 | ~ spl30_21 | ~ spl30_66)), inference(subsumption_resolution, [], [f2457, f213])).
fof(f213, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f2457, plain, ((e0 = e3) | (~ spl30_19 | ~ spl30_21 | ~ spl30_66)), inference(forward_demodulation, [], [f2456, f521])).
fof(f2456, plain, ((e3 = op(e2, e2)) | (~ spl30_19 | ~ spl30_66)), inference(backward_demodulation, [], [f743, f512])).
fof(f743, plain, ((e3 = op(e2, op(e2, e3))) | ~ spl30_66), inference(avatar_component_clause, [], [f742])).
fof(f2437, plain, (spl30_22 | ~ spl30_27 | ~ spl30_68), inference(avatar_contradiction_clause, [], [f2436])).
fof(f2436, plain, ($false | (spl30_22 | ~ spl30_27 | ~ spl30_68)), inference(subsumption_resolution, [], [f2433, f524])).
fof(f524, plain, (~ (e1 = op(e2, e2)) | spl30_22), inference(avatar_component_clause, [], [f523])).
fof(f2433, plain, ((e1 = op(e2, e2)) | (~ spl30_27 | ~ spl30_68)), inference(backward_demodulation, [], [f753, f546])).
fof(f2428, plain, (~ spl30_30 | ~ spl30_32), inference(avatar_contradiction_clause, [], [f2427])).
fof(f2427, plain, ($false | (~ spl30_30 | ~ spl30_32)), inference(subsumption_resolution, [], [f2426, f215])).
fof(f2426, plain, ((e1 = e3) | (~ spl30_30 | ~ spl30_32)), inference(backward_demodulation, [], [f567, f559])).
fof(f2346, plain, (~ spl30_1 | ~ spl30_4), inference(avatar_contradiction_clause, [], [f2345])).
fof(f2345, plain, ($false | (~ spl30_1 | ~ spl30_4)), inference(subsumption_resolution, [], [f2344, f213])).
fof(f2344, plain, ((e0 = e3) | (~ spl30_1 | ~ spl30_4)), inference(forward_demodulation, [], [f448, f436])).
fof(f2327, plain, (~ spl30_24 | ~ spl30_33 | spl30_123 | ~ spl30_124), inference(avatar_contradiction_clause, [], [f2326])).
fof(f2326, plain, ($false | (~ spl30_24 | ~ spl30_33 | spl30_123 | ~ spl30_124)), inference(subsumption_resolution, [], [f2321, f572])).
fof(f2321, plain, (~ (e0 = op(e1, e3)) | (~ spl30_24 | spl30_123 | ~ spl30_124)), inference(backward_demodulation, [], [f2273, f533])).
fof(f2273, plain, (~ (e0 = op(e1, op(e2, e2))) | (spl30_123 | ~ spl30_124)), inference(forward_demodulation, [], [f1176, f1179])).
fof(f1179, plain, ((e1 = op(op(e2, e2), e2)) | ~ spl30_124), inference(avatar_component_clause, [], [f1178])).
fof(f2314, plain, (~ spl30_25 | spl30_30 | ~ spl30_68), inference(avatar_contradiction_clause, [], [f2313])).
fof(f2313, plain, ($false | (~ spl30_25 | spl30_30 | ~ spl30_68)), inference(subsumption_resolution, [], [f2309, f558])).
fof(f558, plain, (~ (e1 = op(e2, e0)) | spl30_30), inference(avatar_component_clause, [], [f557])).
fof(f2309, plain, ((e1 = op(e2, e0)) | (~ spl30_25 | ~ spl30_68)), inference(backward_demodulation, [], [f753, f538])).
fof(f2301, plain, (~ spl30_33 | ~ spl30_48 | spl30_71), inference(avatar_contradiction_clause, [], [f2300])).
fof(f2300, plain, ($false | (~ spl30_33 | ~ spl30_48 | spl30_71)), inference(subsumption_resolution, [], [f2298, f635])).
fof(f2298, plain, (~ (e3 = op(e1, e0)) | (~ spl30_33 | spl30_71)), inference(backward_demodulation, [], [f770, f572])).
fof(f770, plain, (~ (e3 = op(e1, op(e1, e3))) | spl30_71), inference(avatar_component_clause, [], [f768])).
fof(f768, plain, (spl30_71 <=> (e3 = op(e1, op(e1, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl30_71])])).
fof(f2282, plain, (~ spl30_59 | ~ spl30_60), inference(avatar_contradiction_clause, [], [f2281])).
fof(f2281, plain, ($false | (~ spl30_59 | ~ spl30_60)), inference(subsumption_resolution, [], [f2280, f216])).
fof(f2280, plain, ((e2 = e3) | (~ spl30_59 | ~ spl30_60)), inference(forward_demodulation, [], [f686, f682])).
fof(f2277, plain, (spl30_11 | ~ spl30_6 | ~ spl30_82), inference(avatar_split_clause, [], [f2276, f825, f455, f476])).
fof(f2276, plain, ((e2 = op(e3, e1)) | (~ spl30_6 | ~ spl30_82)), inference(forward_demodulation, [], [f826, f457])).
fof(f457, plain, ((e1 = op(e3, e2)) | ~ spl30_6), inference(avatar_component_clause, [], [f455])).
fof(f826, plain, ((e2 = op(e3, op(e3, e2))) | ~ spl30_82), inference(avatar_component_clause, [], [f825])).
fof(f2252, plain, (~ spl30_39 | spl30_72), inference(avatar_contradiction_clause, [], [f2251])).
fof(f2251, plain, ($false | (~ spl30_39 | spl30_72)), inference(subsumption_resolution, [], [f2250, f597])).
fof(f2250, plain, (~ (e2 = op(e1, e2)) | (~ spl30_39 | spl30_72)), inference(forward_demodulation, [], [f775, f597])).
fof(f2244, plain, (~ spl30_33 | ~ spl30_1), inference(avatar_split_clause, [], [f2243, f434, f570])).
fof(f2243, plain, (~ (e0 = op(e1, e3)) | ~ spl30_1), inference(forward_demodulation, [], [f185, f436])).
fof(f2240, plain, (~ spl30_26 | ~ spl30_42), inference(avatar_split_clause, [], [f2239, f608, f540])).
fof(f2239, plain, (~ (e1 = op(e2, e1)) | ~ spl30_42), inference(forward_demodulation, [], [f172, f610])).
fof(f2232, plain, (~ spl30_16 | ~ spl30_1 | spl30_81), inference(avatar_split_clause, [], [f2217, f820, f434, f497])).
fof(f2217, plain, (~ (e3 = op(e3, e0)) | (~ spl30_1 | spl30_81)), inference(backward_demodulation, [], [f822, f436])).
fof(f2225, plain, (~ spl30_1 | ~ spl30_51 | spl30_131), inference(avatar_contradiction_clause, [], [f2224])).
fof(f2224, plain, ($false | (~ spl30_1 | ~ spl30_51 | spl30_131)), inference(subsumption_resolution, [], [f2220, f648])).
fof(f2220, plain, (~ (e2 = op(e0, e3)) | (~ spl30_1 | spl30_131)), inference(backward_demodulation, [], [f1215, f436])).
fof(f2213, plain, (~ spl30_9 | ~ spl30_12), inference(avatar_contradiction_clause, [], [f2212])).
fof(f2212, plain, ($false | (~ spl30_9 | ~ spl30_12)), inference(subsumption_resolution, [], [f2209, f213])).
fof(f2209, plain, ((e0 = e3) | (~ spl30_9 | ~ spl30_12)), inference(backward_demodulation, [], [f482, f470])).
fof(f2199, plain, (~ spl30_21 | ~ spl30_31 | spl30_67), inference(avatar_contradiction_clause, [], [f2198])).
fof(f2198, plain, ($false | (~ spl30_21 | ~ spl30_31 | spl30_67)), inference(subsumption_resolution, [], [f2194, f563])).
fof(f2194, plain, (~ (e2 = op(e2, e0)) | (~ spl30_21 | spl30_67)), inference(backward_demodulation, [], [f749, f521])).
fof(f2190, plain, (~ spl30_31 | ~ spl30_32), inference(avatar_contradiction_clause, [], [f2189])).
fof(f2189, plain, ($false | (~ spl30_31 | ~ spl30_32)), inference(subsumption_resolution, [], [f2188, f216])).
fof(f2188, plain, ((e2 = e3) | (~ spl30_31 | ~ spl30_32)), inference(forward_demodulation, [], [f567, f563])).
fof(f2187, plain, (~ spl30_39 | ~ spl30_40), inference(avatar_contradiction_clause, [], [f2186])).
fof(f2186, plain, ($false | (~ spl30_39 | ~ spl30_40)), inference(subsumption_resolution, [], [f2185, f216])).
fof(f2185, plain, ((e2 = e3) | (~ spl30_39 | ~ spl30_40)), inference(backward_demodulation, [], [f601, f597])).
fof(f2180, plain, (spl30_41 | ~ spl30_42 | ~ spl30_121), inference(avatar_contradiction_clause, [], [f2179])).
fof(f2179, plain, ($false | (spl30_41 | ~ spl30_42 | ~ spl30_121)), inference(subsumption_resolution, [], [f2177, f605])).
fof(f605, plain, (~ (e0 = op(e1, e1)) | spl30_41), inference(avatar_component_clause, [], [f604])).
fof(f2177, plain, ((e0 = op(e1, e1)) | (~ spl30_42 | ~ spl30_121)), inference(backward_demodulation, [], [f1163, f610])).
fof(f2155, plain, (~ spl30_51 | ~ spl30_52), inference(avatar_contradiction_clause, [], [f2154])).
fof(f2154, plain, ($false | (~ spl30_51 | ~ spl30_52)), inference(subsumption_resolution, [], [f2153, f216])).
fof(f2153, plain, ((e2 = e3) | (~ spl30_51 | ~ spl30_52)), inference(backward_demodulation, [], [f652, f648])).
fof(f2137, plain, (~ spl30_17 | ~ spl30_31 | ~ spl30_66), inference(avatar_contradiction_clause, [], [f2136])).
fof(f2136, plain, ($false | (~ spl30_17 | ~ spl30_31 | ~ spl30_66)), inference(subsumption_resolution, [], [f2135, f216])).
fof(f2135, plain, ((e2 = e3) | (~ spl30_17 | ~ spl30_31 | ~ spl30_66)), inference(forward_demodulation, [], [f2134, f563])).
fof(f2134, plain, ((e3 = op(e2, e0)) | (~ spl30_17 | ~ spl30_66)), inference(forward_demodulation, [], [f743, f504])).
fof(f2133, plain, (spl30_21 | ~ spl30_31 | ~ spl30_69), inference(avatar_split_clause, [], [f2132, f757, f561, f519])).
fof(f2132, plain, ((e0 = op(e2, e2)) | (~ spl30_31 | ~ spl30_69)), inference(forward_demodulation, [], [f758, f563])).
fof(f2130, plain, (~ spl30_52 | spl30_76), inference(avatar_contradiction_clause, [], [f2129])).
fof(f2129, plain, ($false | (~ spl30_52 | spl30_76)), inference(subsumption_resolution, [], [f2128, f652])).
fof(f2128, plain, (~ (e3 = op(e0, e3)) | (~ spl30_52 | spl30_76)), inference(forward_demodulation, [], [f796, f652])).
fof(f2093, plain, (~ spl30_13 | spl30_84), inference(avatar_contradiction_clause, [], [f2092])).
fof(f2092, plain, ($false | (~ spl30_13 | spl30_84)), inference(subsumption_resolution, [], [f2088, f487])).
fof(f2088, plain, (~ (e0 = op(e3, e0)) | (~ spl30_13 | spl30_84)), inference(backward_demodulation, [], [f837, f487])).
fof(f2039, plain, (spl30_57 | ~ spl30_62 | ~ spl30_79), inference(avatar_split_clause, [], [f2038, f809, f693, f672])).
fof(f2038, plain, ((e0 = op(e0, e1)) | (~ spl30_62 | ~ spl30_79)), inference(forward_demodulation, [], [f810, f695])).
fof(f2037, plain, (~ spl30_18 | ~ spl30_3 | spl30_130), inference(avatar_split_clause, [], [f2036, f1208, f442, f506])).
fof(f2036, plain, (~ (e1 = op(e2, e3)) | (~ spl30_3 | spl30_130)), inference(forward_demodulation, [], [f1210, f444])).
fof(f2028, plain, (~ spl30_60 | ~ spl30_52), inference(avatar_split_clause, [], [f2027, f650, f684])).
fof(f2027, plain, (~ (e3 = op(e0, e1)) | ~ spl30_52), inference(forward_demodulation, [], [f191, f652])).
fof(f2022, plain, (~ spl30_47 | ~ spl30_62 | spl30_112), inference(avatar_split_clause, [], [f1934, f1118, f693, f629])).
fof(f1934, plain, (~ (e2 = op(e1, e0)) | (~ spl30_62 | spl30_112)), inference(forward_demodulation, [], [f1120, f695])).
fof(f2018, plain, (spl30_44 | ~ spl30_34 | ~ spl30_71), inference(avatar_split_clause, [], [f1976, f768, f574, f616])).
fof(f1976, plain, ((e3 = op(e1, e1)) | (~ spl30_34 | ~ spl30_71)), inference(backward_demodulation, [], [f769, f576])).
fof(f769, plain, ((e3 = op(e1, op(e1, e3))) | ~ spl30_71), inference(avatar_component_clause, [], [f768])).
fof(f2009, plain, (~ spl30_17 | ~ spl30_3 | spl30_133), inference(avatar_split_clause, [], [f1852, f1222, f442, f502])).
fof(f1852, plain, (~ (e0 = op(e2, e3)) | (~ spl30_3 | spl30_133)), inference(forward_demodulation, [], [f1224, f444])).
fof(f1987, plain, (~ spl30_5 | ~ spl30_6), inference(avatar_contradiction_clause, [], [f1986])).
fof(f1986, plain, ($false | (~ spl30_5 | ~ spl30_6)), inference(subsumption_resolution, [], [f1985, f211])).
fof(f211, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f1985, plain, ((e0 = e1) | (~ spl30_5 | ~ spl30_6)), inference(forward_demodulation, [], [f457, f453])).
fof(f1969, plain, (~ spl30_42 | ~ spl30_43), inference(avatar_contradiction_clause, [], [f1968])).
fof(f1968, plain, ($false | (~ spl30_42 | ~ spl30_43)), inference(subsumption_resolution, [], [f1967, f214])).
fof(f1967, plain, ((e1 = e2) | (~ spl30_42 | ~ spl30_43)), inference(backward_demodulation, [], [f614, f610])).
fof(f1966, plain, (~ spl30_49 | ~ spl30_52), inference(avatar_contradiction_clause, [], [f1965])).
fof(f1965, plain, ($false | (~ spl30_49 | ~ spl30_52)), inference(subsumption_resolution, [], [f1964, f213])).
fof(f1964, plain, ((e0 = e3) | (~ spl30_49 | ~ spl30_52)), inference(forward_demodulation, [], [f652, f640])).
fof(f1949, plain, (spl30_38 | ~ spl30_43 | ~ spl30_73), inference(avatar_split_clause, [], [f1911, f778, f612, f591])).
fof(f1911, plain, ((e1 = op(e1, e2)) | (~ spl30_43 | ~ spl30_73)), inference(forward_demodulation, [], [f779, f614])).
fof(f779, plain, ((e1 = op(e1, op(e1, e1))) | ~ spl30_73), inference(avatar_component_clause, [], [f778])).
fof(f1940, plain, (~ spl30_53 | ~ spl30_55), inference(avatar_contradiction_clause, [], [f1939])).
fof(f1939, plain, ($false | (~ spl30_53 | ~ spl30_55)), inference(subsumption_resolution, [], [f1937, f212])).
fof(f1937, plain, ((e0 = e2) | (~ spl30_53 | ~ spl30_55)), inference(backward_demodulation, [], [f665, f657])).
fof(f1931, plain, (~ spl30_53 | ~ spl30_5), inference(avatar_split_clause, [], [f1557, f451, f655])).
fof(f1557, plain, (~ (e0 = op(e0, e2)) | ~ spl30_5), inference(forward_demodulation, [], [f177, f453])).
fof(f1900, plain, (~ spl30_41 | ~ spl30_43), inference(avatar_contradiction_clause, [], [f1899])).
fof(f1899, plain, ($false | (~ spl30_41 | ~ spl30_43)), inference(subsumption_resolution, [], [f1898, f212])).
fof(f1898, plain, ((e0 = e2) | (~ spl30_41 | ~ spl30_43)), inference(forward_demodulation, [], [f614, f606])).
fof(f1877, plain, (~ spl30_58 | ~ spl30_62), inference(avatar_split_clause, [], [f1868, f693, f676])).
fof(f1868, plain, (~ (e1 = op(e0, e1)) | ~ spl30_62), inference(backward_demodulation, [], [f187, f695])).
fof(f1875, plain, (~ spl30_30 | ~ spl30_62), inference(avatar_split_clause, [], [f1866, f693, f557])).
fof(f1866, plain, (~ (e1 = op(e2, e0)) | ~ spl30_62), inference(backward_demodulation, [], [f164, f695])).
fof(f1860, plain, (~ spl30_41 | spl30_46 | ~ spl30_73), inference(avatar_contradiction_clause, [], [f1859])).
fof(f1859, plain, ($false | (~ spl30_41 | spl30_46 | ~ spl30_73)), inference(subsumption_resolution, [], [f1858, f626])).
fof(f626, plain, (~ (e1 = op(e1, e0)) | spl30_46), inference(avatar_component_clause, [], [f625])).
fof(f1858, plain, ((e1 = op(e1, e0)) | (~ spl30_41 | ~ spl30_73)), inference(forward_demodulation, [], [f779, f606])).
fof(f1857, plain, (~ spl30_55 | spl30_77), inference(avatar_contradiction_clause, [], [f1856])).
fof(f1856, plain, ($false | (~ spl30_55 | spl30_77)), inference(subsumption_resolution, [], [f1855, f665])).
fof(f1855, plain, (~ (e2 = op(e0, e2)) | (~ spl30_55 | spl30_77)), inference(forward_demodulation, [], [f801, f665])).
fof(f801, plain, (~ (e2 = op(e0, op(e0, e2))) | spl30_77), inference(avatar_component_clause, [], [f799])).
fof(f799, plain, (spl30_77 <=> (e2 = op(e0, op(e0, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl30_77])])).
fof(f1850, plain, (~ spl30_64 | ~ spl30_49 | spl30_76), inference(avatar_split_clause, [], [f1778, f794, f638, f701])).
fof(f1778, plain, (~ (op(e0, e0) = e3) | (~ spl30_49 | spl30_76)), inference(backward_demodulation, [], [f796, f640])).
fof(f1847, plain, (~ spl30_60 | ~ spl30_41 | spl30_119), inference(avatar_split_clause, [], [f1790, f1153, f604, f684])).
fof(f1790, plain, (~ (e3 = op(e0, e1)) | (~ spl30_41 | spl30_119)), inference(backward_demodulation, [], [f1155, f606])).
fof(f1827, plain, (~ spl30_13 | ~ spl30_14), inference(avatar_contradiction_clause, [], [f1826])).
fof(f1826, plain, ($false | (~ spl30_13 | ~ spl30_14)), inference(subsumption_resolution, [], [f1825, f211])).
fof(f1825, plain, ((e0 = e1) | (~ spl30_13 | ~ spl30_14)), inference(backward_demodulation, [], [f491, f487])).
fof(f1804, plain, (~ spl30_25 | ~ spl30_27), inference(avatar_contradiction_clause, [], [f1803])).
fof(f1803, plain, ($false | (~ spl30_25 | ~ spl30_27)), inference(subsumption_resolution, [], [f1802, f212])).
fof(f1802, plain, ((e0 = e2) | (~ spl30_25 | ~ spl30_27)), inference(backward_demodulation, [], [f546, f538])).
fof(f1793, plain, (~ spl30_33 | ~ spl30_41), inference(avatar_split_clause, [], [f1785, f604, f570])).
fof(f1785, plain, (~ (e0 = op(e1, e3)) | ~ spl30_41), inference(backward_demodulation, [], [f197, f606])).
fof(f1774, plain, (~ spl30_58 | spl30_78), inference(avatar_contradiction_clause, [], [f1773])).
fof(f1773, plain, ($false | (~ spl30_58 | spl30_78)), inference(subsumption_resolution, [], [f1769, f678])).
fof(f1769, plain, (~ (e1 = op(e0, e1)) | (~ spl30_58 | spl30_78)), inference(backward_demodulation, [], [f806, f678])).
fof(f1772, plain, (~ spl30_50 | ~ spl30_58), inference(avatar_split_clause, [], [f1768, f676, f642])).
fof(f1768, plain, (~ (e1 = op(e0, e3)) | ~ spl30_58), inference(backward_demodulation, [], [f191, f678])).
fof(f1753, plain, (spl30_49 | ~ spl30_64 | ~ spl30_79), inference(avatar_split_clause, [], [f1676, f809, f701, f638])).
fof(f1676, plain, ((e0 = op(e0, e3)) | (~ spl30_64 | ~ spl30_79)), inference(backward_demodulation, [], [f810, f703])).
fof(f1722, plain, (~ spl30_20 | spl30_66), inference(avatar_contradiction_clause, [], [f1721])).
fof(f1721, plain, ($false | (~ spl30_20 | spl30_66)), inference(subsumption_resolution, [], [f1718, f516])).
fof(f1718, plain, (~ (e3 = op(e2, e3)) | (~ spl30_20 | spl30_66)), inference(backward_demodulation, [], [f744, f516])).
fof(f1717, plain, (~ spl30_21 | ~ spl30_24), inference(avatar_contradiction_clause, [], [f1716])).
fof(f1716, plain, ($false | (~ spl30_21 | ~ spl30_24)), inference(subsumption_resolution, [], [f1715, f213])).
fof(f1715, plain, ((e0 = e3) | (~ spl30_21 | ~ spl30_24)), inference(backward_demodulation, [], [f533, f521])).
fof(f1691, plain, (~ spl30_45 | ~ spl30_47), inference(avatar_contradiction_clause, [], [f1690])).
fof(f1690, plain, ($false | (~ spl30_45 | ~ spl30_47)), inference(subsumption_resolution, [], [f1689, f212])).
fof(f1689, plain, ((e0 = e2) | (~ spl30_45 | ~ spl30_47)), inference(forward_demodulation, [], [f631, f623])).
fof(f1656, plain, (~ spl30_28 | ~ spl30_24), inference(avatar_split_clause, [], [f1655, f531, f548])).
fof(f1655, plain, (~ (e3 = op(e2, e1)) | ~ spl30_24), inference(forward_demodulation, [], [f202, f533])).
fof(f1648, plain, (~ spl30_11 | ~ spl30_12), inference(avatar_contradiction_clause, [], [f1647])).
fof(f1647, plain, ($false | (~ spl30_11 | ~ spl30_12)), inference(subsumption_resolution, [], [f1646, f216])).
fof(f1646, plain, ((e2 = e3) | (~ spl30_11 | ~ spl30_12)), inference(backward_demodulation, [], [f482, f478])).
fof(f1640, plain, (~ spl30_14 | ~ spl30_16), inference(avatar_contradiction_clause, [], [f1639])).
fof(f1639, plain, ($false | (~ spl30_14 | ~ spl30_16)), inference(subsumption_resolution, [], [f1638, f215])).
fof(f1638, plain, ((e1 = e3) | (~ spl30_14 | ~ spl30_16)), inference(backward_demodulation, [], [f499, f491])).
fof(f1634, plain, (~ spl30_5 | ~ spl30_24 | spl30_127), inference(avatar_contradiction_clause, [], [f1633])).
fof(f1633, plain, ($false | (~ spl30_5 | ~ spl30_24 | spl30_127)), inference(subsumption_resolution, [], [f1628, f453])).
fof(f1628, plain, (~ (e0 = op(e3, e2)) | (~ spl30_24 | spl30_127)), inference(backward_demodulation, [], [f1194, f533])).
fof(f1624, plain, (~ spl30_26 | spl30_68), inference(avatar_contradiction_clause, [], [f1623])).
fof(f1623, plain, ($false | (~ spl30_26 | spl30_68)), inference(subsumption_resolution, [], [f1620, f542])).
fof(f1620, plain, (~ (e1 = op(e2, e1)) | (~ spl30_26 | spl30_68)), inference(backward_demodulation, [], [f754, f542])).
fof(f1613, plain, (~ spl30_23 | ~ spl30_31), inference(avatar_split_clause, [], [f1606, f561, f527])).
fof(f1606, plain, (~ (e2 = op(e2, e2)) | ~ spl30_31), inference(backward_demodulation, [], [f200, f563])).
fof(f1604, plain, (~ spl30_38 | ~ spl30_39), inference(avatar_contradiction_clause, [], [f1603])).
fof(f1603, plain, ($false | (~ spl30_38 | ~ spl30_39)), inference(subsumption_resolution, [], [f1602, f214])).
fof(f1602, plain, ((e1 = e2) | (~ spl30_38 | ~ spl30_39)), inference(forward_demodulation, [], [f597, f593])).
fof(f1599, plain, (~ spl30_28 | ~ spl30_44), inference(avatar_split_clause, [], [f1593, f616, f548])).
fof(f1593, plain, (~ (e3 = op(e2, e1)) | ~ spl30_44), inference(backward_demodulation, [], [f172, f618])).
fof(f1581, plain, (~ spl30_57 | ~ spl30_60), inference(avatar_contradiction_clause, [], [f1580])).
fof(f1580, plain, ($false | (~ spl30_57 | ~ spl30_60)), inference(subsumption_resolution, [], [f1579, f213])).
fof(f1579, plain, ((e0 = e3) | (~ spl30_57 | ~ spl30_60)), inference(forward_demodulation, [], [f686, f674])).
fof(f1559, plain, (spl30_62 | ~ spl30_57 | ~ spl30_78), inference(avatar_split_clause, [], [f1457, f804, f672, f693])).
fof(f1457, plain, ((op(e0, e0) = e1) | (~ spl30_57 | ~ spl30_78)), inference(forward_demodulation, [], [f805, f674])).
fof(f805, plain, ((e1 = op(e0, op(e0, e1))) | ~ spl30_78), inference(avatar_component_clause, [], [f804])).
fof(f1547, plain, (~ spl30_29 | ~ spl30_45), inference(avatar_split_clause, [], [f1546, f621, f553])).
fof(f1546, plain, (~ (e0 = op(e2, e0)) | ~ spl30_45), inference(forward_demodulation, [], [f166, f623])).
fof(f1534, plain, (~ spl30_2 | ~ spl30_3), inference(avatar_contradiction_clause, [], [f1533])).
fof(f1533, plain, ($false | (~ spl30_2 | ~ spl30_3)), inference(subsumption_resolution, [], [f1532, f214])).
fof(f1532, plain, ((e1 = e2) | (~ spl30_2 | ~ spl30_3)), inference(backward_demodulation, [], [f444, f440])).
fof(f1524, plain, (~ spl30_15 | ~ spl30_16), inference(avatar_contradiction_clause, [], [f1523])).
fof(f1523, plain, ($false | (~ spl30_15 | ~ spl30_16)), inference(subsumption_resolution, [], [f1522, f216])).
fof(f1522, plain, ((e2 = e3) | (~ spl30_15 | ~ spl30_16)), inference(backward_demodulation, [], [f499, f495])).
fof(f1511, plain, (~ spl30_22 | ~ spl30_23), inference(avatar_contradiction_clause, [], [f1510])).
fof(f1510, plain, ($false | (~ spl30_22 | ~ spl30_23)), inference(subsumption_resolution, [], [f1509, f214])).
fof(f1509, plain, ((e1 = e2) | (~ spl30_22 | ~ spl30_23)), inference(forward_demodulation, [], [f529, f525])).
fof(f1508, plain, (~ spl30_27 | ~ spl30_28), inference(avatar_contradiction_clause, [], [f1507])).
fof(f1507, plain, ($false | (~ spl30_27 | ~ spl30_28)), inference(subsumption_resolution, [], [f1506, f216])).
fof(f1506, plain, ((e2 = e3) | (~ spl30_27 | ~ spl30_28)), inference(backward_demodulation, [], [f550, f546])).
fof(f1503, plain, (~ spl30_35 | ~ spl30_36), inference(avatar_contradiction_clause, [], [f1502])).
fof(f1502, plain, ($false | (~ spl30_35 | ~ spl30_36)), inference(subsumption_resolution, [], [f1501, f216])).
fof(f1501, plain, ((e2 = e3) | (~ spl30_35 | ~ spl30_36)), inference(backward_demodulation, [], [f584, f580])).
fof(f1490, plain, (~ spl30_42 | spl30_73), inference(avatar_contradiction_clause, [], [f1489])).
fof(f1489, plain, ($false | (~ spl30_42 | spl30_73)), inference(subsumption_resolution, [], [f1485, f610])).
fof(f1485, plain, (~ (e1 = op(e1, e1)) | (~ spl30_42 | spl30_73)), inference(backward_demodulation, [], [f780, f610])).
fof(f1479, plain, (~ spl30_45 | spl30_74), inference(avatar_contradiction_clause, [], [f1478])).
fof(f1478, plain, ($false | (~ spl30_45 | spl30_74)), inference(subsumption_resolution, [], [f1473, f623])).
fof(f1473, plain, (~ (e0 = op(e1, e0)) | (~ spl30_45 | spl30_74)), inference(backward_demodulation, [], [f785, f623])).
fof(f1467, plain, (~ spl30_51 | ~ spl30_56 | spl30_76), inference(avatar_contradiction_clause, [], [f1466])).
fof(f1466, plain, ($false | (~ spl30_51 | ~ spl30_56 | spl30_76)), inference(subsumption_resolution, [], [f1464, f669])).
fof(f1464, plain, (~ (e3 = op(e0, e2)) | (~ spl30_51 | spl30_76)), inference(backward_demodulation, [], [f796, f648])).
fof(f1450, plain, (~ spl30_51 | ~ spl30_3), inference(avatar_split_clause, [], [f1449, f442, f646])).
fof(f1449, plain, (~ (e2 = op(e0, e3)) | ~ spl30_3), inference(forward_demodulation, [], [f183, f444])).
fof(f1427, plain, (~ spl30_29 | spl30_69), inference(avatar_contradiction_clause, [], [f1426])).
fof(f1426, plain, ($false | (~ spl30_29 | spl30_69)), inference(subsumption_resolution, [], [f1425, f555])).
fof(f1425, plain, (~ (e0 = op(e2, e0)) | (~ spl30_29 | spl30_69)), inference(forward_demodulation, [], [f759, f555])).
fof(f1424, plain, (~ spl30_27 | ~ spl30_22 | spl30_67), inference(avatar_split_clause, [], [f1360, f747, f523, f544])).
fof(f1360, plain, (~ (e2 = op(e2, e1)) | (~ spl30_22 | spl30_67)), inference(backward_demodulation, [], [f749, f525])).
fof(f1413, plain, (~ spl30_10 | spl30_83), inference(avatar_contradiction_clause, [], [f1412])).
fof(f1412, plain, ($false | (~ spl30_10 | spl30_83)), inference(subsumption_resolution, [], [f1411, f474])).
fof(f1411, plain, (~ (e1 = op(e3, e1)) | (~ spl30_10 | spl30_83)), inference(forward_demodulation, [], [f832, f474])).
fof(f1408, plain, (~ spl30_8 | ~ spl30_3 | spl30_81), inference(avatar_split_clause, [], [f1396, f820, f442, f463])).
fof(f1396, plain, (~ (e3 = op(e3, e2)) | (~ spl30_3 | spl30_81)), inference(backward_demodulation, [], [f822, f444])).
fof(f1407, plain, (~ spl30_1 | ~ spl30_3), inference(avatar_contradiction_clause, [], [f1406])).
fof(f1406, plain, ($false | (~ spl30_1 | ~ spl30_3)), inference(subsumption_resolution, [], [f1405, f212])).
fof(f1405, plain, ((e0 = e2) | (~ spl30_1 | ~ spl30_3)), inference(backward_demodulation, [], [f444, f436])).
fof(f1394, plain, (~ spl30_7 | spl30_82), inference(avatar_contradiction_clause, [], [f1393])).
fof(f1393, plain, ($false | (~ spl30_7 | spl30_82)), inference(subsumption_resolution, [], [f1391, f461])).
fof(f1391, plain, (~ (e2 = op(e3, e2)) | (~ spl30_7 | spl30_82)), inference(backward_demodulation, [], [f827, f461])).
fof(f1388, plain, (~ spl30_10 | ~ spl30_12), inference(avatar_contradiction_clause, [], [f1387])).
fof(f1387, plain, ($false | (~ spl30_10 | ~ spl30_12)), inference(subsumption_resolution, [], [f1386, f215])).
fof(f1386, plain, ((e1 = e3) | (~ spl30_10 | ~ spl30_12)), inference(backward_demodulation, [], [f482, f474])).
fof(f1355, plain, (~ spl30_18 | ~ spl30_28 | spl30_68), inference(avatar_split_clause, [], [f1350, f752, f548, f506])).
fof(f1350, plain, (~ (e1 = op(e2, e3)) | (~ spl30_28 | spl30_68)), inference(backward_demodulation, [], [f754, f550])).
fof(f1346, plain, (~ spl30_29 | ~ spl30_32), inference(avatar_contradiction_clause, [], [f1345])).
fof(f1345, plain, ($false | (~ spl30_29 | ~ spl30_32)), inference(subsumption_resolution, [], [f1344, f213])).
fof(f1344, plain, ((e0 = e3) | (~ spl30_29 | ~ spl30_32)), inference(backward_demodulation, [], [f567, f555])).
fof(f1333, plain, (~ spl30_34 | ~ spl30_36), inference(avatar_contradiction_clause, [], [f1332])).
fof(f1332, plain, ($false | (~ spl30_34 | ~ spl30_36)), inference(subsumption_resolution, [], [f1331, f215])).
fof(f1331, plain, ((e1 = e3) | (~ spl30_34 | ~ spl30_36)), inference(backward_demodulation, [], [f584, f576])).
fof(f1330, plain, (~ spl30_36 | spl30_71), inference(avatar_contradiction_clause, [], [f1329])).
fof(f1329, plain, ($false | (~ spl30_36 | spl30_71)), inference(subsumption_resolution, [], [f1326, f584])).
fof(f1326, plain, (~ (e3 = op(e1, e3)) | (~ spl30_36 | spl30_71)), inference(backward_demodulation, [], [f770, f584])).
fof(f1328, plain, (~ spl30_4 | ~ spl30_36), inference(avatar_split_clause, [], [f1325, f582, f446])).
fof(f1325, plain, (~ (e3 = op(e3, e3)) | ~ spl30_36), inference(backward_demodulation, [], [f185, f584])).
fof(f1320, plain, (~ spl30_5 | ~ spl30_37), inference(avatar_split_clause, [], [f1316, f587, f451])).
fof(f1316, plain, (~ (e0 = op(e3, e2)) | ~ spl30_37), inference(backward_demodulation, [], [f179, f589])).
fof(f1313, plain, (~ spl30_38 | ~ spl30_43 | spl30_73), inference(avatar_split_clause, [], [f1306, f778, f612, f591])).
fof(f1306, plain, (~ (e1 = op(e1, e2)) | (~ spl30_43 | spl30_73)), inference(backward_demodulation, [], [f780, f614])).
fof(f1300, plain, (~ spl30_28 | ~ spl30_47 | ~ spl30_62 | spl30_116), inference(avatar_split_clause, [], [f1292, f1138, f693, f629, f548])).
fof(f1292, plain, (~ (e3 = op(e2, e1)) | (~ spl30_47 | ~ spl30_62 | spl30_116)), inference(backward_demodulation, [], [f1243, f631])).
fof(f1243, plain, (~ (e3 = op(op(e1, e0), e1)) | (~ spl30_62 | spl30_116)), inference(backward_demodulation, [], [f1140, f695])).
fof(f1276, plain, (~ spl30_55 | ~ spl30_56), inference(avatar_contradiction_clause, [], [f1275])).
fof(f1275, plain, ($false | (~ spl30_55 | ~ spl30_56)), inference(subsumption_resolution, [], [f1274, f216])).
fof(f1274, plain, ((e2 = e3) | (~ spl30_55 | ~ spl30_56)), inference(backward_demodulation, [], [f669, f665])).
fof(f1263, plain, (~ spl30_57 | ~ spl30_62 | spl30_78), inference(avatar_contradiction_clause, [], [f1262])).
fof(f1262, plain, ($false | (~ spl30_57 | ~ spl30_62 | spl30_78)), inference(subsumption_resolution, [], [f1256, f695])).
fof(f1256, plain, (~ (op(e0, e0) = e1) | (~ spl30_57 | spl30_78)), inference(backward_demodulation, [], [f806, f674])).
fof(f1232, plain, (~ spl30_134 | ~ spl30_130 | ~ spl30_1), inference(avatar_split_clause, [], [f432, f434, f1208, f1228])).
fof(f432, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (~ (e0 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ~ ((e0 = op(e3, e3)) & (e1 = op(op(e3, e3), e3)) & (e2 = op(op(op(e3, e3), e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax29)).
fof(f1231, plain, (~ spl30_134 | ~ spl30_133 | ~ spl30_2), inference(avatar_split_clause, [], [f431, f438, f1222, f1228])).
fof(f431, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e1 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e2 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e1 = op(e3, e3)) & (e0 = op(op(e3, e3), e3)) & (e2 = op(op(op(e3, e3), e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax28)).
fof(f1226, plain, (~ spl30_132 | ~ spl30_131 | ~ spl30_1), inference(avatar_split_clause, [], [f430, f434, f1213, f1218])).
fof(f430, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (~ (e0 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e1 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((e0 = op(e3, e3)) & (e2 = op(op(e3, e3), e3)) & (e1 = op(op(op(e3, e3), e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax27)).
fof(f1225, plain, (~ spl30_132 | ~ spl30_133 | ~ spl30_3), inference(avatar_split_clause, [], [f429, f442, f1222, f1218])).
fof(f429, plain, (~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e2 = op(e3, e3)) | ~ (e0 = op(op(e3, e3), e3)) | ~ (e1 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(e3, e3)) & (e0 = op(op(e3, e3), e3)) & (e1 = op(op(op(e3, e3), e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax26)).
fof(f1216, plain, (~ spl30_129 | ~ spl30_131 | ~ spl30_2), inference(avatar_split_clause, [], [f428, f438, f1213, f1204])).
fof(f428, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(cnf_transformation, [], [f49])).
fof(f49, plain, (~ (e1 = op(e3, e3)) | ~ (e2 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ ((e1 = op(e3, e3)) & (e2 = op(op(e3, e3), e3)) & (e0 = op(op(op(e3, e3), e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax25)).
fof(f1211, plain, (~ spl30_129 | ~ spl30_130 | ~ spl30_3), inference(avatar_split_clause, [], [f427, f442, f1208, f1204])).
fof(f427, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(cnf_transformation, [], [f48])).
fof(f48, plain, (~ (e2 = op(e3, e3)) | ~ (e1 = op(op(e3, e3), e3)) | ~ (e0 = op(op(op(e3, e3), e3), op(e3, e3)))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ~ ((e2 = op(e3, e3)) & (e1 = op(op(e3, e3), e3)) & (e0 = op(op(op(e3, e3), e3), op(e3, e3)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax24)).
fof(f1202, plain, (~ spl30_128 | ~ spl30_124 | ~ spl30_21), inference(avatar_split_clause, [], [f426, f519, f1178, f1198])).
fof(f426, plain, (~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e0 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(e2, e2)) & (e1 = op(op(e2, e2), e2)) & (e3 = op(op(op(e2, e2), e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax23)).
fof(f1201, plain, (~ spl30_128 | ~ spl30_127 | ~ spl30_22), inference(avatar_split_clause, [], [f425, f523, f1192, f1198])).
fof(f425, plain, (~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e1 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e3 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(e2, e2)) & (e0 = op(op(e2, e2), e2)) & (e3 = op(op(op(e2, e2), e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax22)).
fof(f1196, plain, (~ spl30_126 | ~ spl30_125 | ~ spl30_21), inference(avatar_split_clause, [], [f424, f519, f1183, f1188])).
fof(f424, plain, (~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(e2, e2)) & (e3 = op(op(e2, e2), e2)) & (e1 = op(op(op(e2, e2), e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax21)).
fof(f1195, plain, (~ spl30_126 | ~ spl30_127 | ~ spl30_24), inference(avatar_split_clause, [], [f423, f531, f1192, f1188])).
fof(f423, plain, (~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e3 = op(e2, e2)) | ~ (e0 = op(op(e2, e2), e2)) | ~ (e1 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e3 = op(e2, e2)) & (e0 = op(op(e2, e2), e2)) & (e1 = op(op(op(e2, e2), e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax20)).
fof(f1186, plain, (~ spl30_123 | ~ spl30_125 | ~ spl30_22), inference(avatar_split_clause, [], [f422, f523, f1183, f1174])).
fof(f422, plain, (~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e1 = op(e2, e2)) | ~ (e3 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(e2, e2)) & (e3 = op(op(e2, e2), e2)) & (e0 = op(op(op(e2, e2), e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax19)).
fof(f1181, plain, (~ spl30_123 | ~ spl30_124 | ~ spl30_24), inference(avatar_split_clause, [], [f421, f531, f1178, f1174])).
fof(f421, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(e2, e2)) | ~ (e1 = op(op(e2, e2), e2)) | ~ (e0 = op(op(op(e2, e2), e2), op(e2, e2)))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(e2, e2)) & (e1 = op(op(e2, e2), e2)) & (e0 = op(op(op(e2, e2), e2), op(e2, e2)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax18)).
fof(f1172, plain, (~ spl30_122 | ~ spl30_118 | ~ spl30_41), inference(avatar_split_clause, [], [f420, f604, f1148, f1168])).
fof(f420, plain, (~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e0 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(e1, e1)) & (e2 = op(op(e1, e1), e1)) & (e3 = op(op(op(e1, e1), e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax17)).
fof(f1171, plain, (~ spl30_122 | ~ spl30_121 | ~ spl30_43), inference(avatar_split_clause, [], [f419, f612, f1162, f1168])).
fof(f419, plain, (~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (~ (e2 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ~ ((e2 = op(e1, e1)) & (e0 = op(op(e1, e1), e1)) & (e3 = op(op(op(e1, e1), e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax16)).
fof(f1166, plain, (~ spl30_120 | ~ spl30_119 | ~ spl30_41), inference(avatar_split_clause, [], [f418, f604, f1153, f1158])).
fof(f418, plain, (~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e0 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(e1, e1)) & (e3 = op(op(e1, e1), e1)) & (e2 = op(op(op(e1, e1), e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax15)).
fof(f1165, plain, (~ spl30_120 | ~ spl30_121 | ~ spl30_44), inference(avatar_split_clause, [], [f417, f616, f1162, f1158])).
fof(f417, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (~ (e3 = op(e1, e1)) | ~ (e0 = op(op(e1, e1), e1)) | ~ (e2 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ~ ((e3 = op(e1, e1)) & (e0 = op(op(e1, e1), e1)) & (e2 = op(op(op(e1, e1), e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax14)).
fof(f1156, plain, (~ spl30_117 | ~ spl30_119 | ~ spl30_43), inference(avatar_split_clause, [], [f416, f612, f1153, f1144])).
fof(f416, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(cnf_transformation, [], [f37])).
fof(f37, plain, (~ (e2 = op(e1, e1)) | ~ (e3 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ((e2 = op(e1, e1)) & (e3 = op(op(e1, e1), e1)) & (e0 = op(op(op(e1, e1), e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax13)).
fof(f1151, plain, (~ spl30_117 | ~ spl30_118 | ~ spl30_44), inference(avatar_split_clause, [], [f415, f616, f1148, f1144])).
fof(f415, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e3 = op(e1, e1)) | ~ (e2 = op(op(e1, e1), e1)) | ~ (e0 = op(op(op(e1, e1), e1), op(e1, e1)))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(e1, e1)) & (e2 = op(op(e1, e1), e1)) & (e0 = op(op(op(e1, e1), e1), op(e1, e1)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax12)).
fof(f1142, plain, (~ spl30_116 | ~ spl30_112 | ~ spl30_62), inference(avatar_split_clause, [], [f414, f693, f1118, f1138])).
fof(f414, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (op(e0, e0) = e1) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((op(e0, e0) = e1) & (e2 = op(op(e0, e0), e0)) & (e3 = op(op(op(e0, e0), e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax11)).
fof(f1141, plain, (~ spl30_116 | ~ spl30_115 | ~ spl30_63), inference(avatar_split_clause, [], [f413, f697, f1132, f1138])).
fof(f413, plain, (~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (op(e0, e0) = e2) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e3 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((op(e0, e0) = e2) & (e1 = op(op(e0, e0), e0)) & (e3 = op(op(op(e0, e0), e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax10)).
fof(f1136, plain, (~ spl30_114 | ~ spl30_113 | ~ spl30_62), inference(avatar_split_clause, [], [f412, f693, f1123, f1128])).
fof(f412, plain, (~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (op(e0, e0) = e1) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((op(e0, e0) = e1) & (e3 = op(op(e0, e0), e0)) & (e2 = op(op(op(e0, e0), e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax9)).
fof(f1135, plain, (~ spl30_114 | ~ spl30_115 | ~ spl30_64), inference(avatar_split_clause, [], [f411, f701, f1132, f1128])).
fof(f411, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (op(e0, e0) = e3) | ~ (e1 = op(op(e0, e0), e0)) | ~ (e2 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((op(e0, e0) = e3) & (e1 = op(op(e0, e0), e0)) & (e2 = op(op(op(e0, e0), e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax8)).
fof(f1126, plain, (~ spl30_111 | ~ spl30_113 | ~ spl30_63), inference(avatar_split_clause, [], [f410, f697, f1123, f1114])).
fof(f410, plain, (~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (op(e0, e0) = e2) | ~ (e3 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((op(e0, e0) = e2) & (e3 = op(op(e0, e0), e0)) & (e1 = op(op(op(e0, e0), e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax7)).
fof(f1121, plain, (~ spl30_111 | ~ spl30_112 | ~ spl30_64), inference(avatar_split_clause, [], [f409, f701, f1118, f1114])).
fof(f409, plain, (~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (op(e0, e0) = e3) | ~ (e2 = op(op(e0, e0), e0)) | ~ (e1 = op(op(op(e0, e0), e0), op(e0, e0)))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((op(e0, e0) = e3) & (e2 = op(op(e0, e0), e0)) & (e1 = op(op(op(e0, e0), e0), op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax6)).
fof(f1112, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | spl30_4), inference(avatar_split_clause, [], [f397, f446, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f1092, plain, (spl30_110 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl30_110])])).
fof(f1082, plain, (spl30_109 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl30_109])])).
fof(f1072, plain, (spl30_108 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl30_108])])).
fof(f1062, plain, (spl30_107 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl30_107])])).
fof(f1052, plain, (spl30_106 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl30_106])])).
fof(f1042, plain, (spl30_105 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl30_105])])).
fof(f1032, plain, (spl30_104 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl30_104])])).
fof(f1022, plain, (spl30_103 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl30_103])])).
fof(f1012, plain, (spl30_102 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl30_102])])).
fof(f1002, plain, (spl30_101 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl30_101])])).
fof(f992, plain, (spl30_100 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl30_100])])).
fof(f982, plain, (spl30_99 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl30_99])])).
fof(f972, plain, (spl30_98 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl30_98])])).
fof(f962, plain, (spl30_97 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl30_97])])).
fof(f952, plain, (spl30_96 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl30_96])])).
fof(f397, plain, ((e3 = op(e3, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f84, plain, (((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15) & ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e0 = op(e0, e1)) & (op(e0, e0) = e1)) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e0 = op(e0, e1)) & (op(e0, e0) = e1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e0 = op(e0, e2)) & (op(e0, e0) = e2)) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e0 = op(e0, e2)) & (op(e0, e0) = e2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e0 = op(e0, e3)) & (op(e0, e0) = e3)) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e0 = op(e0, e3)) & (op(e0, e0) = e3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e1 = op(e1, e0)) & (e0 = op(e1, e1))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e1 = op(e1, e2)) & (e2 = op(e1, e1))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e1 = op(e1, e3)) & (e3 = op(e1, e1))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e2 = op(e2, e0)) & (e0 = op(e2, e2))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e2 = op(e2, e1)) & (e1 = op(e2, e2))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e2 = op(e2, e3)) & (e3 = op(e2, e2))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e3 = op(e3, e0)) & (e0 = op(e3, e3))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e3 = op(e3, e1)) & (e1 = op(e3, e3))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e3 = op(e3, e2)) & (e2 = op(e3, e3))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f5, plain, (((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e3 = op(e3, e2)) & (e2 = op(e3, e3))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e3 = op(e3, e1)) & (e1 = op(e3, e3))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e3 = op(e3, e0)) & (e0 = op(e3, e3))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e2 = op(e2, e3)) & (e3 = op(e2, e2))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e2 = op(e2, e1)) & (e1 = op(e2, e2))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e2 = op(e2, e0)) & (e0 = op(e2, e2))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e1 = op(e1, e3)) & (e3 = op(e1, e1))) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e1 = op(e1, e2)) & (e2 = op(e1, e1))) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e1 = op(e1, e0)) & (e0 = op(e1, e1))) | (~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e0 = op(e0, e3)) & (op(e0, e0) = e3)) | (~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e0 = op(e0, e2)) & (op(e0, e0) = e2)) | (~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e0 = op(e0, e1)) & (op(e0, e0) = e1)) | (~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax5)).
fof(f1111, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | ~ spl30_4), inference(avatar_split_clause, [], [f398, f446, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f398, plain, (~ (e3 = op(e3, e3)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1110, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | ~ spl30_84), inference(avatar_split_clause, [], [f399, f835, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f399, plain, (~ (e0 = op(e3, op(e3, e0))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1109, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | ~ spl30_83), inference(avatar_split_clause, [], [f400, f830, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f400, plain, (~ (e1 = op(e3, op(e3, e1))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1108, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | ~ spl30_82), inference(avatar_split_clause, [], [f401, f825, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f401, plain, (~ (e2 = op(e3, op(e3, e2))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1107, plain, (spl30_110 | spl30_109 | spl30_108 | spl30_107 | spl30_106 | spl30_105 | spl30_104 | spl30_103 | spl30_102 | spl30_101 | spl30_100 | spl30_99 | spl30_98 | spl30_97 | spl30_96 | ~ spl30_81), inference(avatar_split_clause, [], [f402, f820, f952, f962, f972, f982, f992, f1002, f1012, f1022, f1032, f1042, f1052, f1062, f1072, f1082, f1092])).
fof(f402, plain, (~ (e3 = op(e3, op(e3, e3))) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f84])).
fof(f1106, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_92 | spl30_91 | spl30_90 | spl30_89 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_75 | spl30_70 | spl30_65 | spl30_4), inference(avatar_split_clause, [], [f403, f446, f738, f764, f790, f816, f842, f852, f862, f872, f882, f892, f902, f912, f922, f932, f942])).
fof(f942, plain, (spl30_95 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl30_95])])).
fof(f932, plain, (spl30_94 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl30_94])])).
fof(f922, plain, (spl30_93 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl30_93])])).
fof(f912, plain, (spl30_92 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl30_92])])).
fof(f902, plain, (spl30_91 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl30_91])])).
fof(f892, plain, (spl30_90 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl30_90])])).
fof(f882, plain, (spl30_89 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl30_89])])).
fof(f872, plain, (spl30_88 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl30_88])])).
fof(f862, plain, (spl30_87 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl30_87])])).
fof(f852, plain, (spl30_86 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl30_86])])).
fof(f842, plain, (spl30_85 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl30_85])])).
fof(f816, plain, (spl30_80 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl30_80])])).
fof(f790, plain, (spl30_75 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl30_75])])).
fof(f764, plain, (spl30_70 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl30_70])])).
fof(f738, plain, (spl30_65 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl30_65])])).
fof(f403, plain, ((e3 = op(e3, e3)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1105, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_92 | spl30_91 | spl30_90 | spl30_89 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_75 | spl30_70 | spl30_65 | ~ spl30_4), inference(avatar_split_clause, [], [f404, f446, f738, f764, f790, f816, f842, f852, f862, f872, f882, f892, f902, f912, f922, f932, f942])).
fof(f404, plain, (~ (e3 = op(e3, e3)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1104, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_92 | spl30_91 | spl30_90 | spl30_89 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_75 | spl30_70 | spl30_65 | ~ spl30_84), inference(avatar_split_clause, [], [f405, f835, f738, f764, f790, f816, f842, f852, f862, f872, f882, f892, f902, f912, f922, f932, f942])).
fof(f405, plain, (~ (e0 = op(e3, op(e3, e0))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1103, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_92 | spl30_91 | spl30_90 | spl30_89 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_75 | spl30_70 | spl30_65 | ~ spl30_83), inference(avatar_split_clause, [], [f406, f830, f738, f764, f790, f816, f842, f852, f862, f872, f882, f892, f902, f912, f922, f932, f942])).
fof(f406, plain, (~ (e1 = op(e3, op(e3, e1))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1102, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_92 | spl30_91 | spl30_90 | spl30_89 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_75 | spl30_70 | spl30_65 | ~ spl30_82), inference(avatar_split_clause, [], [f407, f825, f738, f764, f790, f816, f842, f852, f862, f872, f882, f892, f902, f912, f922, f932, f942])).
fof(f407, plain, (~ (e2 = op(e3, op(e3, e2))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1101, plain, (spl30_95 | spl30_94 | spl30_93 | spl30_92 | spl30_91 | spl30_90 | spl30_89 | spl30_88 | spl30_87 | spl30_86 | spl30_85 | spl30_80 | spl30_75 | spl30_70 | spl30_65 | ~ spl30_81), inference(avatar_split_clause, [], [f408, f820, f738, f764, f790, f816, f842, f852, f862, f872, f882, f892, f902, f912, f922, f932, f942])).
fof(f408, plain, (~ (e3 = op(e3, op(e3, e3))) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15), inference(cnf_transformation, [], [f84])).
fof(f1100, plain, (~ spl30_110 | spl30_61), inference(avatar_split_clause, [], [f391, f689, f1092])).
fof(f391, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1099, plain, (~ spl30_110 | ~ spl30_61), inference(avatar_split_clause, [], [f392, f689, f1092])).
fof(f392, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f114])).
fof(f1090, plain, (~ spl30_109 | spl30_62), inference(avatar_split_clause, [], [f385, f693, f1082])).
fof(f385, plain, ((op(e0, e0) = e1) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e0 = op(e0, e1)) & (op(e0, e0) = e1)) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1089, plain, (~ spl30_109 | ~ spl30_57), inference(avatar_split_clause, [], [f386, f672, f1082])).
fof(f386, plain, (~ (e0 = op(e0, e1)) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1088, plain, (~ spl30_109 | ~ spl30_74), inference(avatar_split_clause, [], [f387, f783, f1082])).
fof(f387, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1087, plain, (~ spl30_109 | ~ spl30_73), inference(avatar_split_clause, [], [f388, f778, f1082])).
fof(f388, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1086, plain, (~ spl30_109 | ~ spl30_72), inference(avatar_split_clause, [], [f389, f773, f1082])).
fof(f389, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1085, plain, (~ spl30_109 | ~ spl30_71), inference(avatar_split_clause, [], [f390, f768, f1082])).
fof(f390, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f1080, plain, (~ spl30_108 | spl30_63), inference(avatar_split_clause, [], [f379, f697, f1072])).
fof(f379, plain, ((op(e0, e0) = e2) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e0 = op(e0, e2)) & (op(e0, e0) = e2)) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1079, plain, (~ spl30_108 | ~ spl30_53), inference(avatar_split_clause, [], [f380, f655, f1072])).
fof(f380, plain, (~ (e0 = op(e0, e2)) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1078, plain, (~ spl30_108 | ~ spl30_69), inference(avatar_split_clause, [], [f381, f757, f1072])).
fof(f381, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1077, plain, (~ spl30_108 | ~ spl30_68), inference(avatar_split_clause, [], [f382, f752, f1072])).
fof(f382, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1076, plain, (~ spl30_108 | ~ spl30_67), inference(avatar_split_clause, [], [f383, f747, f1072])).
fof(f383, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1075, plain, (~ spl30_108 | ~ spl30_66), inference(avatar_split_clause, [], [f384, f742, f1072])).
fof(f384, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP2), inference(cnf_transformation, [], [f112])).
fof(f1070, plain, (~ spl30_107 | spl30_64), inference(avatar_split_clause, [], [f373, f701, f1062])).
fof(f373, plain, ((op(e0, e0) = e3) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e0 = op(e0, e3)) & (op(e0, e0) = e3)) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1069, plain, (~ spl30_107 | ~ spl30_49), inference(avatar_split_clause, [], [f374, f638, f1062])).
fof(f374, plain, (~ (e0 = op(e0, e3)) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1068, plain, (~ spl30_107 | ~ spl30_84), inference(avatar_split_clause, [], [f375, f835, f1062])).
fof(f375, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1067, plain, (~ spl30_107 | ~ spl30_83), inference(avatar_split_clause, [], [f376, f830, f1062])).
fof(f376, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1066, plain, (~ spl30_107 | ~ spl30_82), inference(avatar_split_clause, [], [f377, f825, f1062])).
fof(f377, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1065, plain, (~ spl30_107 | ~ spl30_81), inference(avatar_split_clause, [], [f378, f820, f1062])).
fof(f378, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP3), inference(cnf_transformation, [], [f111])).
fof(f1060, plain, (~ spl30_106 | spl30_41), inference(avatar_split_clause, [], [f367, f604, f1052])).
fof(f367, plain, ((e0 = op(e1, e1)) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e1 = op(e1, e0)) & (e0 = op(e1, e1))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1059, plain, (~ spl30_106 | ~ spl30_46), inference(avatar_split_clause, [], [f368, f625, f1052])).
fof(f368, plain, (~ (e1 = op(e1, e0)) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1058, plain, (~ spl30_106 | ~ spl30_79), inference(avatar_split_clause, [], [f369, f809, f1052])).
fof(f369, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1057, plain, (~ spl30_106 | ~ spl30_78), inference(avatar_split_clause, [], [f370, f804, f1052])).
fof(f370, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1056, plain, (~ spl30_106 | ~ spl30_77), inference(avatar_split_clause, [], [f371, f799, f1052])).
fof(f371, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1055, plain, (~ spl30_106 | ~ spl30_76), inference(avatar_split_clause, [], [f372, f794, f1052])).
fof(f372, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP4), inference(cnf_transformation, [], [f110])).
fof(f1050, plain, (~ spl30_105 | spl30_42), inference(avatar_split_clause, [], [f361, f608, f1042])).
fof(f361, plain, ((e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1049, plain, (~ spl30_105 | ~ spl30_42), inference(avatar_split_clause, [], [f362, f608, f1042])).
fof(f362, plain, (~ (e1 = op(e1, e1)) | ~ sP5), inference(cnf_transformation, [], [f109])).
fof(f1040, plain, (~ spl30_104 | spl30_43), inference(avatar_split_clause, [], [f355, f612, f1032])).
fof(f355, plain, ((e2 = op(e1, e1)) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e1 = op(e1, e2)) & (e2 = op(e1, e1))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1039, plain, (~ spl30_104 | ~ spl30_38), inference(avatar_split_clause, [], [f356, f591, f1032])).
fof(f356, plain, (~ (e1 = op(e1, e2)) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1038, plain, (~ spl30_104 | ~ spl30_69), inference(avatar_split_clause, [], [f357, f757, f1032])).
fof(f357, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1037, plain, (~ spl30_104 | ~ spl30_68), inference(avatar_split_clause, [], [f358, f752, f1032])).
fof(f358, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1036, plain, (~ spl30_104 | ~ spl30_67), inference(avatar_split_clause, [], [f359, f747, f1032])).
fof(f359, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1035, plain, (~ spl30_104 | ~ spl30_66), inference(avatar_split_clause, [], [f360, f742, f1032])).
fof(f360, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP6), inference(cnf_transformation, [], [f108])).
fof(f1030, plain, (~ spl30_103 | spl30_44), inference(avatar_split_clause, [], [f349, f616, f1022])).
fof(f349, plain, ((e3 = op(e1, e1)) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e1 = op(e1, e3)) & (e3 = op(e1, e1))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1029, plain, (~ spl30_103 | ~ spl30_34), inference(avatar_split_clause, [], [f350, f574, f1022])).
fof(f350, plain, (~ (e1 = op(e1, e3)) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1028, plain, (~ spl30_103 | ~ spl30_84), inference(avatar_split_clause, [], [f351, f835, f1022])).
fof(f351, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1027, plain, (~ spl30_103 | ~ spl30_83), inference(avatar_split_clause, [], [f352, f830, f1022])).
fof(f352, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1026, plain, (~ spl30_103 | ~ spl30_82), inference(avatar_split_clause, [], [f353, f825, f1022])).
fof(f353, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1025, plain, (~ spl30_103 | ~ spl30_81), inference(avatar_split_clause, [], [f354, f820, f1022])).
fof(f354, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP7), inference(cnf_transformation, [], [f107])).
fof(f1020, plain, (~ spl30_102 | spl30_21), inference(avatar_split_clause, [], [f343, f519, f1012])).
fof(f343, plain, ((e0 = op(e2, e2)) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e2 = op(e2, e0)) & (e0 = op(e2, e2))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1019, plain, (~ spl30_102 | ~ spl30_31), inference(avatar_split_clause, [], [f344, f561, f1012])).
fof(f344, plain, (~ (e2 = op(e2, e0)) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1018, plain, (~ spl30_102 | ~ spl30_79), inference(avatar_split_clause, [], [f345, f809, f1012])).
fof(f345, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1017, plain, (~ spl30_102 | ~ spl30_78), inference(avatar_split_clause, [], [f346, f804, f1012])).
fof(f346, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1016, plain, (~ spl30_102 | ~ spl30_77), inference(avatar_split_clause, [], [f347, f799, f1012])).
fof(f347, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1015, plain, (~ spl30_102 | ~ spl30_76), inference(avatar_split_clause, [], [f348, f794, f1012])).
fof(f348, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP8), inference(cnf_transformation, [], [f106])).
fof(f1010, plain, (~ spl30_101 | spl30_22), inference(avatar_split_clause, [], [f337, f523, f1002])).
fof(f337, plain, ((e1 = op(e2, e2)) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e2 = op(e2, e1)) & (e1 = op(e2, e2))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1009, plain, (~ spl30_101 | ~ spl30_27), inference(avatar_split_clause, [], [f338, f544, f1002])).
fof(f338, plain, (~ (e2 = op(e2, e1)) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1008, plain, (~ spl30_101 | ~ spl30_74), inference(avatar_split_clause, [], [f339, f783, f1002])).
fof(f339, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1007, plain, (~ spl30_101 | ~ spl30_73), inference(avatar_split_clause, [], [f340, f778, f1002])).
fof(f340, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1006, plain, (~ spl30_101 | ~ spl30_72), inference(avatar_split_clause, [], [f341, f773, f1002])).
fof(f341, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1005, plain, (~ spl30_101 | ~ spl30_71), inference(avatar_split_clause, [], [f342, f768, f1002])).
fof(f342, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP9), inference(cnf_transformation, [], [f105])).
fof(f1000, plain, (~ spl30_100 | spl30_23), inference(avatar_split_clause, [], [f331, f527, f992])).
fof(f331, plain, ((e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f999, plain, (~ spl30_100 | ~ spl30_23), inference(avatar_split_clause, [], [f332, f527, f992])).
fof(f332, plain, (~ (e2 = op(e2, e2)) | ~ sP10), inference(cnf_transformation, [], [f104])).
fof(f990, plain, (~ spl30_99 | spl30_24), inference(avatar_split_clause, [], [f325, f531, f982])).
fof(f325, plain, ((e3 = op(e2, e2)) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e2 = op(e2, e3)) & (e3 = op(e2, e2))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f989, plain, (~ spl30_99 | ~ spl30_19), inference(avatar_split_clause, [], [f326, f510, f982])).
fof(f326, plain, (~ (e2 = op(e2, e3)) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f988, plain, (~ spl30_99 | ~ spl30_84), inference(avatar_split_clause, [], [f327, f835, f982])).
fof(f327, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f987, plain, (~ spl30_99 | ~ spl30_83), inference(avatar_split_clause, [], [f328, f830, f982])).
fof(f328, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f986, plain, (~ spl30_99 | ~ spl30_82), inference(avatar_split_clause, [], [f329, f825, f982])).
fof(f329, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f985, plain, (~ spl30_99 | ~ spl30_81), inference(avatar_split_clause, [], [f330, f820, f982])).
fof(f330, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP11), inference(cnf_transformation, [], [f103])).
fof(f980, plain, (~ spl30_98 | spl30_1), inference(avatar_split_clause, [], [f319, f434, f972])).
fof(f319, plain, ((e0 = op(e3, e3)) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e3 = op(e3, e0)) & (e0 = op(e3, e3))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f979, plain, (~ spl30_98 | ~ spl30_16), inference(avatar_split_clause, [], [f320, f497, f972])).
fof(f320, plain, (~ (e3 = op(e3, e0)) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f978, plain, (~ spl30_98 | ~ spl30_79), inference(avatar_split_clause, [], [f321, f809, f972])).
fof(f321, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f977, plain, (~ spl30_98 | ~ spl30_78), inference(avatar_split_clause, [], [f322, f804, f972])).
fof(f322, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f976, plain, (~ spl30_98 | ~ spl30_77), inference(avatar_split_clause, [], [f323, f799, f972])).
fof(f323, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f975, plain, (~ spl30_98 | ~ spl30_76), inference(avatar_split_clause, [], [f324, f794, f972])).
fof(f324, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP12), inference(cnf_transformation, [], [f102])).
fof(f970, plain, (~ spl30_97 | spl30_2), inference(avatar_split_clause, [], [f313, f438, f962])).
fof(f313, plain, ((e1 = op(e3, e3)) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e3 = op(e3, e1)) & (e1 = op(e3, e3))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f969, plain, (~ spl30_97 | ~ spl30_12), inference(avatar_split_clause, [], [f314, f480, f962])).
fof(f314, plain, (~ (e3 = op(e3, e1)) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f968, plain, (~ spl30_97 | ~ spl30_74), inference(avatar_split_clause, [], [f315, f783, f962])).
fof(f315, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f967, plain, (~ spl30_97 | ~ spl30_73), inference(avatar_split_clause, [], [f316, f778, f962])).
fof(f316, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f966, plain, (~ spl30_97 | ~ spl30_72), inference(avatar_split_clause, [], [f317, f773, f962])).
fof(f317, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f965, plain, (~ spl30_97 | ~ spl30_71), inference(avatar_split_clause, [], [f318, f768, f962])).
fof(f318, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP13), inference(cnf_transformation, [], [f101])).
fof(f960, plain, (~ spl30_96 | spl30_3), inference(avatar_split_clause, [], [f307, f442, f952])).
fof(f307, plain, ((e2 = op(e3, e3)) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e3 = op(e3, e2)) & (e2 = op(e3, e3))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f959, plain, (~ spl30_96 | ~ spl30_8), inference(avatar_split_clause, [], [f308, f463, f952])).
fof(f308, plain, (~ (e3 = op(e3, e2)) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f958, plain, (~ spl30_96 | ~ spl30_69), inference(avatar_split_clause, [], [f309, f757, f952])).
fof(f309, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f957, plain, (~ spl30_96 | ~ spl30_68), inference(avatar_split_clause, [], [f310, f752, f952])).
fof(f310, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f956, plain, (~ spl30_96 | ~ spl30_67), inference(avatar_split_clause, [], [f311, f747, f952])).
fof(f311, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f955, plain, (~ spl30_96 | ~ spl30_66), inference(avatar_split_clause, [], [f312, f742, f952])).
fof(f312, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP14), inference(cnf_transformation, [], [f100])).
fof(f950, plain, (~ spl30_95 | spl30_61), inference(avatar_split_clause, [], [f301, f689, f942])).
fof(f301, plain, ((e0 = op(e0, e0)) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f949, plain, (~ spl30_95 | ~ spl30_61), inference(avatar_split_clause, [], [f302, f689, f942])).
fof(f302, plain, (~ (e0 = op(e0, e0)) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f946, plain, (~ spl30_95 | ~ spl30_77), inference(avatar_split_clause, [], [f305, f799, f942])).
fof(f305, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f945, plain, (~ spl30_95 | ~ spl30_76), inference(avatar_split_clause, [], [f306, f794, f942])).
fof(f306, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP15), inference(cnf_transformation, [], [f99])).
fof(f940, plain, (~ spl30_94 | spl30_57), inference(avatar_split_clause, [], [f295, f672, f932])).
fof(f295, plain, ((e0 = op(e0, e1)) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f939, plain, (~ spl30_94 | ~ spl30_62), inference(avatar_split_clause, [], [f296, f693, f932])).
fof(f296, plain, (~ (op(e0, e0) = e1) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f938, plain, (~ spl30_94 | ~ spl30_74), inference(avatar_split_clause, [], [f297, f783, f932])).
fof(f297, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f937, plain, (~ spl30_94 | ~ spl30_73), inference(avatar_split_clause, [], [f298, f778, f932])).
fof(f298, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f936, plain, (~ spl30_94 | ~ spl30_72), inference(avatar_split_clause, [], [f299, f773, f932])).
fof(f299, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f935, plain, (~ spl30_94 | ~ spl30_71), inference(avatar_split_clause, [], [f300, f768, f932])).
fof(f300, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP16), inference(cnf_transformation, [], [f98])).
fof(f930, plain, (~ spl30_93 | spl30_53), inference(avatar_split_clause, [], [f289, f655, f922])).
fof(f289, plain, ((e0 = op(e0, e2)) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f929, plain, (~ spl30_93 | ~ spl30_63), inference(avatar_split_clause, [], [f290, f697, f922])).
fof(f290, plain, (~ (op(e0, e0) = e2) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f928, plain, (~ spl30_93 | ~ spl30_69), inference(avatar_split_clause, [], [f291, f757, f922])).
fof(f291, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f927, plain, (~ spl30_93 | ~ spl30_68), inference(avatar_split_clause, [], [f292, f752, f922])).
fof(f292, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f926, plain, (~ spl30_93 | ~ spl30_67), inference(avatar_split_clause, [], [f293, f747, f922])).
fof(f293, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f925, plain, (~ spl30_93 | ~ spl30_66), inference(avatar_split_clause, [], [f294, f742, f922])).
fof(f294, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP17), inference(cnf_transformation, [], [f97])).
fof(f920, plain, (~ spl30_92 | spl30_49), inference(avatar_split_clause, [], [f283, f638, f912])).
fof(f283, plain, ((e0 = op(e0, e3)) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f919, plain, (~ spl30_92 | ~ spl30_64), inference(avatar_split_clause, [], [f284, f701, f912])).
fof(f284, plain, (~ (op(e0, e0) = e3) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f918, plain, (~ spl30_92 | ~ spl30_84), inference(avatar_split_clause, [], [f285, f835, f912])).
fof(f285, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f917, plain, (~ spl30_92 | ~ spl30_83), inference(avatar_split_clause, [], [f286, f830, f912])).
fof(f286, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f916, plain, (~ spl30_92 | ~ spl30_82), inference(avatar_split_clause, [], [f287, f825, f912])).
fof(f287, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f915, plain, (~ spl30_92 | ~ spl30_81), inference(avatar_split_clause, [], [f288, f820, f912])).
fof(f288, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP18), inference(cnf_transformation, [], [f96])).
fof(f910, plain, (~ spl30_91 | spl30_46), inference(avatar_split_clause, [], [f277, f625, f902])).
fof(f277, plain, ((e1 = op(e1, e0)) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f909, plain, (~ spl30_91 | ~ spl30_41), inference(avatar_split_clause, [], [f278, f604, f902])).
fof(f278, plain, (~ (e0 = op(e1, e1)) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f908, plain, (~ spl30_91 | ~ spl30_79), inference(avatar_split_clause, [], [f279, f809, f902])).
fof(f279, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f907, plain, (~ spl30_91 | ~ spl30_78), inference(avatar_split_clause, [], [f280, f804, f902])).
fof(f280, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f906, plain, (~ spl30_91 | ~ spl30_77), inference(avatar_split_clause, [], [f281, f799, f902])).
fof(f281, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f905, plain, (~ spl30_91 | ~ spl30_76), inference(avatar_split_clause, [], [f282, f794, f902])).
fof(f282, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP19), inference(cnf_transformation, [], [f95])).
fof(f900, plain, (~ spl30_90 | spl30_42), inference(avatar_split_clause, [], [f271, f608, f892])).
fof(f271, plain, ((e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f899, plain, (~ spl30_90 | ~ spl30_42), inference(avatar_split_clause, [], [f272, f608, f892])).
fof(f272, plain, (~ (e1 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f898, plain, (~ spl30_90 | ~ spl30_74), inference(avatar_split_clause, [], [f273, f783, f892])).
fof(f273, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP20), inference(cnf_transformation, [], [f94])).
fof(f890, plain, (~ spl30_89 | spl30_38), inference(avatar_split_clause, [], [f265, f591, f882])).
fof(f265, plain, ((e1 = op(e1, e2)) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f889, plain, (~ spl30_89 | ~ spl30_43), inference(avatar_split_clause, [], [f266, f612, f882])).
fof(f266, plain, (~ (e2 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f888, plain, (~ spl30_89 | ~ spl30_69), inference(avatar_split_clause, [], [f267, f757, f882])).
fof(f267, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f887, plain, (~ spl30_89 | ~ spl30_68), inference(avatar_split_clause, [], [f268, f752, f882])).
fof(f268, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f886, plain, (~ spl30_89 | ~ spl30_67), inference(avatar_split_clause, [], [f269, f747, f882])).
fof(f269, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f885, plain, (~ spl30_89 | ~ spl30_66), inference(avatar_split_clause, [], [f270, f742, f882])).
fof(f270, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP21), inference(cnf_transformation, [], [f93])).
fof(f880, plain, (~ spl30_88 | spl30_34), inference(avatar_split_clause, [], [f259, f574, f872])).
fof(f259, plain, ((e1 = op(e1, e3)) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f879, plain, (~ spl30_88 | ~ spl30_44), inference(avatar_split_clause, [], [f260, f616, f872])).
fof(f260, plain, (~ (e3 = op(e1, e1)) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f878, plain, (~ spl30_88 | ~ spl30_84), inference(avatar_split_clause, [], [f261, f835, f872])).
fof(f261, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f877, plain, (~ spl30_88 | ~ spl30_83), inference(avatar_split_clause, [], [f262, f830, f872])).
fof(f262, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f876, plain, (~ spl30_88 | ~ spl30_82), inference(avatar_split_clause, [], [f263, f825, f872])).
fof(f263, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f875, plain, (~ spl30_88 | ~ spl30_81), inference(avatar_split_clause, [], [f264, f820, f872])).
fof(f264, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP22), inference(cnf_transformation, [], [f92])).
fof(f870, plain, (~ spl30_87 | spl30_31), inference(avatar_split_clause, [], [f253, f561, f862])).
fof(f253, plain, ((e2 = op(e2, e0)) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f869, plain, (~ spl30_87 | ~ spl30_21), inference(avatar_split_clause, [], [f254, f519, f862])).
fof(f254, plain, (~ (e0 = op(e2, e2)) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f868, plain, (~ spl30_87 | ~ spl30_79), inference(avatar_split_clause, [], [f255, f809, f862])).
fof(f255, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f867, plain, (~ spl30_87 | ~ spl30_78), inference(avatar_split_clause, [], [f256, f804, f862])).
fof(f256, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f866, plain, (~ spl30_87 | ~ spl30_77), inference(avatar_split_clause, [], [f257, f799, f862])).
fof(f257, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f865, plain, (~ spl30_87 | ~ spl30_76), inference(avatar_split_clause, [], [f258, f794, f862])).
fof(f258, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP23), inference(cnf_transformation, [], [f91])).
fof(f860, plain, (~ spl30_86 | spl30_27), inference(avatar_split_clause, [], [f247, f544, f852])).
fof(f247, plain, ((e2 = op(e2, e1)) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f859, plain, (~ spl30_86 | ~ spl30_22), inference(avatar_split_clause, [], [f248, f523, f852])).
fof(f248, plain, (~ (e1 = op(e2, e2)) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f858, plain, (~ spl30_86 | ~ spl30_74), inference(avatar_split_clause, [], [f249, f783, f852])).
fof(f249, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f857, plain, (~ spl30_86 | ~ spl30_73), inference(avatar_split_clause, [], [f250, f778, f852])).
fof(f250, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f856, plain, (~ spl30_86 | ~ spl30_72), inference(avatar_split_clause, [], [f251, f773, f852])).
fof(f251, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f855, plain, (~ spl30_86 | ~ spl30_71), inference(avatar_split_clause, [], [f252, f768, f852])).
fof(f252, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP24), inference(cnf_transformation, [], [f90])).
fof(f850, plain, (~ spl30_85 | spl30_23), inference(avatar_split_clause, [], [f241, f527, f842])).
fof(f241, plain, ((e2 = op(e2, e2)) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f849, plain, (~ spl30_85 | ~ spl30_23), inference(avatar_split_clause, [], [f242, f527, f842])).
fof(f242, plain, (~ (e2 = op(e2, e2)) | ~ sP25), inference(cnf_transformation, [], [f89])).
fof(f840, plain, (~ spl30_80 | spl30_19), inference(avatar_split_clause, [], [f235, f510, f816])).
fof(f235, plain, ((e2 = op(e2, e3)) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e3 = op(e3, op(e3, e3))) & ~ (e2 = op(e3, op(e3, e2))) & ~ (e1 = op(e3, op(e3, e1))) & ~ (e0 = op(e3, op(e3, e0))) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f839, plain, (~ spl30_80 | ~ spl30_24), inference(avatar_split_clause, [], [f236, f531, f816])).
fof(f236, plain, (~ (e3 = op(e2, e2)) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f838, plain, (~ spl30_80 | ~ spl30_84), inference(avatar_split_clause, [], [f237, f835, f816])).
fof(f237, plain, (~ (e0 = op(e3, op(e3, e0))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f833, plain, (~ spl30_80 | ~ spl30_83), inference(avatar_split_clause, [], [f238, f830, f816])).
fof(f238, plain, (~ (e1 = op(e3, op(e3, e1))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f828, plain, (~ spl30_80 | ~ spl30_82), inference(avatar_split_clause, [], [f239, f825, f816])).
fof(f239, plain, (~ (e2 = op(e3, op(e3, e2))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f823, plain, (~ spl30_80 | ~ spl30_81), inference(avatar_split_clause, [], [f240, f820, f816])).
fof(f240, plain, (~ (e3 = op(e3, op(e3, e3))) | ~ sP26), inference(cnf_transformation, [], [f88])).
fof(f814, plain, (~ spl30_75 | spl30_16), inference(avatar_split_clause, [], [f229, f497, f790])).
fof(f229, plain, ((e3 = op(e3, e0)) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e3 = op(e0, op(e0, e3))) & ~ (e2 = op(e0, op(e0, e2))) & ~ (e1 = op(e0, op(e0, e1))) & ~ (e0 = op(e0, op(e0, e0))) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f813, plain, (~ spl30_75 | ~ spl30_1), inference(avatar_split_clause, [], [f230, f434, f790])).
fof(f230, plain, (~ (e0 = op(e3, e3)) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f812, plain, (~ spl30_75 | ~ spl30_79), inference(avatar_split_clause, [], [f231, f809, f790])).
fof(f231, plain, (~ (e0 = op(e0, op(e0, e0))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f807, plain, (~ spl30_75 | ~ spl30_78), inference(avatar_split_clause, [], [f232, f804, f790])).
fof(f232, plain, (~ (e1 = op(e0, op(e0, e1))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f802, plain, (~ spl30_75 | ~ spl30_77), inference(avatar_split_clause, [], [f233, f799, f790])).
fof(f233, plain, (~ (e2 = op(e0, op(e0, e2))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f797, plain, (~ spl30_75 | ~ spl30_76), inference(avatar_split_clause, [], [f234, f794, f790])).
fof(f234, plain, (~ (e3 = op(e0, op(e0, e3))) | ~ sP27), inference(cnf_transformation, [], [f87])).
fof(f788, plain, (~ spl30_70 | spl30_12), inference(avatar_split_clause, [], [f223, f480, f764])).
fof(f223, plain, ((e3 = op(e3, e1)) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e3 = op(e1, op(e1, e3))) & ~ (e2 = op(e1, op(e1, e2))) & ~ (e1 = op(e1, op(e1, e1))) & ~ (e0 = op(e1, op(e1, e0))) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f787, plain, (~ spl30_70 | ~ spl30_2), inference(avatar_split_clause, [], [f224, f438, f764])).
fof(f224, plain, (~ (e1 = op(e3, e3)) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f786, plain, (~ spl30_70 | ~ spl30_74), inference(avatar_split_clause, [], [f225, f783, f764])).
fof(f225, plain, (~ (e0 = op(e1, op(e1, e0))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f781, plain, (~ spl30_70 | ~ spl30_73), inference(avatar_split_clause, [], [f226, f778, f764])).
fof(f226, plain, (~ (e1 = op(e1, op(e1, e1))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f776, plain, (~ spl30_70 | ~ spl30_72), inference(avatar_split_clause, [], [f227, f773, f764])).
fof(f227, plain, (~ (e2 = op(e1, op(e1, e2))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f771, plain, (~ spl30_70 | ~ spl30_71), inference(avatar_split_clause, [], [f228, f768, f764])).
fof(f228, plain, (~ (e3 = op(e1, op(e1, e3))) | ~ sP28), inference(cnf_transformation, [], [f86])).
fof(f762, plain, (~ spl30_65 | spl30_8), inference(avatar_split_clause, [], [f217, f463, f738])).
fof(f217, plain, ((e3 = op(e3, e2)) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e3 = op(e2, op(e2, e3))) & ~ (e2 = op(e2, op(e2, e2))) & ~ (e1 = op(e2, op(e2, e1))) & ~ (e0 = op(e2, op(e2, e0))) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f761, plain, (~ spl30_65 | ~ spl30_3), inference(avatar_split_clause, [], [f218, f442, f738])).
fof(f218, plain, (~ (e2 = op(e3, e3)) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f760, plain, (~ spl30_65 | ~ spl30_69), inference(avatar_split_clause, [], [f219, f757, f738])).
fof(f219, plain, (~ (e0 = op(e2, op(e2, e0))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f755, plain, (~ spl30_65 | ~ spl30_68), inference(avatar_split_clause, [], [f220, f752, f738])).
fof(f220, plain, (~ (e1 = op(e2, op(e2, e1))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f750, plain, (~ spl30_65 | ~ spl30_67), inference(avatar_split_clause, [], [f221, f747, f738])).
fof(f221, plain, (~ (e2 = op(e2, op(e2, e2))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f745, plain, (~ spl30_65 | ~ spl30_66), inference(avatar_split_clause, [], [f222, f742, f738])).
fof(f222, plain, (~ (e3 = op(e2, op(e2, e3))) | ~ sP29), inference(cnf_transformation, [], [f85])).
fof(f736, plain, (spl30_61 | spl30_57 | spl30_53 | spl30_49), inference(avatar_split_clause, [], [f131, f638, f655, f672, f689])).
fof(f131, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax2)).
fof(f735, plain, (spl30_61 | spl30_45 | spl30_29 | spl30_13), inference(avatar_split_clause, [], [f132, f485, f553, f621, f689])).
fof(f132, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f734, plain, (spl30_62 | spl30_58 | spl30_54 | spl30_50), inference(avatar_split_clause, [], [f133, f642, f659, f676, f693])).
fof(f133, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f733, plain, (spl30_62 | spl30_46 | spl30_30 | spl30_14), inference(avatar_split_clause, [], [f134, f489, f557, f625, f693])).
fof(f134, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f732, plain, (spl30_63 | spl30_59 | spl30_55 | spl30_51), inference(avatar_split_clause, [], [f135, f646, f663, f680, f697])).
fof(f135, plain, ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f731, plain, (spl30_63 | spl30_47 | spl30_31 | spl30_15), inference(avatar_split_clause, [], [f136, f493, f561, f629, f697])).
fof(f136, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f730, plain, (spl30_64 | spl30_60 | spl30_56 | spl30_52), inference(avatar_split_clause, [], [f137, f650, f667, f684, f701])).
fof(f137, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f729, plain, (spl30_64 | spl30_48 | spl30_32 | spl30_16), inference(avatar_split_clause, [], [f138, f497, f565, f633, f701])).
fof(f138, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f728, plain, (spl30_45 | spl30_41 | spl30_37 | spl30_33), inference(avatar_split_clause, [], [f139, f570, f587, f604, f621])).
fof(f139, plain, ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f727, plain, (spl30_57 | spl30_41 | spl30_25 | spl30_9), inference(avatar_split_clause, [], [f140, f468, f536, f604, f672])).
fof(f140, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f726, plain, (spl30_46 | spl30_42 | spl30_38 | spl30_34), inference(avatar_split_clause, [], [f141, f574, f591, f608, f625])).
fof(f141, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f725, plain, (spl30_58 | spl30_42 | spl30_26 | spl30_10), inference(avatar_split_clause, [], [f142, f472, f540, f608, f676])).
fof(f142, plain, ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f724, plain, (spl30_47 | spl30_43 | spl30_39 | spl30_35), inference(avatar_split_clause, [], [f143, f578, f595, f612, f629])).
fof(f143, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f723, plain, (spl30_59 | spl30_43 | spl30_27 | spl30_11), inference(avatar_split_clause, [], [f144, f476, f544, f612, f680])).
fof(f144, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f722, plain, (spl30_48 | spl30_44 | spl30_40 | spl30_36), inference(avatar_split_clause, [], [f145, f582, f599, f616, f633])).
fof(f145, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f721, plain, (spl30_60 | spl30_44 | spl30_28 | spl30_12), inference(avatar_split_clause, [], [f146, f480, f548, f616, f684])).
fof(f146, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f720, plain, (spl30_29 | spl30_25 | spl30_21 | spl30_17), inference(avatar_split_clause, [], [f147, f502, f519, f536, f553])).
fof(f147, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f719, plain, (spl30_53 | spl30_37 | spl30_21 | spl30_5), inference(avatar_split_clause, [], [f148, f451, f519, f587, f655])).
fof(f148, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f718, plain, (spl30_30 | spl30_26 | spl30_22 | spl30_18), inference(avatar_split_clause, [], [f149, f506, f523, f540, f557])).
fof(f149, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f717, plain, (spl30_54 | spl30_38 | spl30_22 | spl30_6), inference(avatar_split_clause, [], [f150, f455, f523, f591, f659])).
fof(f150, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f716, plain, (spl30_31 | spl30_27 | spl30_23 | spl30_19), inference(avatar_split_clause, [], [f151, f510, f527, f544, f561])).
fof(f151, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f715, plain, (spl30_55 | spl30_39 | spl30_23 | spl30_7), inference(avatar_split_clause, [], [f152, f459, f527, f595, f663])).
fof(f152, plain, ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f714, plain, (spl30_32 | spl30_28 | spl30_24 | spl30_20), inference(avatar_split_clause, [], [f153, f514, f531, f548, f565])).
fof(f153, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f713, plain, (spl30_56 | spl30_40 | spl30_24 | spl30_8), inference(avatar_split_clause, [], [f154, f463, f531, f599, f667])).
fof(f154, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f712, plain, (spl30_13 | spl30_9 | spl30_5 | spl30_1), inference(avatar_split_clause, [], [f155, f434, f451, f468, f485])).
fof(f155, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f711, plain, (spl30_49 | spl30_33 | spl30_17 | spl30_1), inference(avatar_split_clause, [], [f156, f434, f502, f570, f638])).
fof(f156, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f710, plain, (spl30_14 | spl30_10 | spl30_6 | spl30_2), inference(avatar_split_clause, [], [f157, f438, f455, f472, f489])).
fof(f157, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f709, plain, (spl30_50 | spl30_34 | spl30_18 | spl30_2), inference(avatar_split_clause, [], [f158, f438, f506, f574, f642])).
fof(f158, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f708, plain, (spl30_15 | spl30_11 | spl30_7 | spl30_3), inference(avatar_split_clause, [], [f159, f442, f459, f476, f493])).
fof(f159, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f707, plain, (spl30_51 | spl30_35 | spl30_19 | spl30_3), inference(avatar_split_clause, [], [f160, f442, f510, f578, f646])).
fof(f160, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f706, plain, (spl30_16 | spl30_12 | spl30_8 | spl30_4), inference(avatar_split_clause, [], [f161, f446, f463, f480, f497])).
fof(f161, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f705, plain, (spl30_52 | spl30_36 | spl30_20 | spl30_4), inference(avatar_split_clause, [], [f162, f446, f514, f582, f650])).
fof(f162, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f704, plain, (spl30_61 | spl30_62 | spl30_63 | spl30_64), inference(avatar_split_clause, [], [f115, f701, f697, f693, f689])).
fof(f115, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG163+1.p', ax1)).
fof(f687, plain, (spl30_57 | spl30_58 | spl30_59 | spl30_60), inference(avatar_split_clause, [], [f116, f684, f680, f676, f672])).
fof(f116, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f670, plain, (spl30_53 | spl30_54 | spl30_55 | spl30_56), inference(avatar_split_clause, [], [f117, f667, f663, f659, f655])).
fof(f117, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f653, plain, (spl30_49 | spl30_50 | spl30_51 | spl30_52), inference(avatar_split_clause, [], [f118, f650, f646, f642, f638])).
fof(f118, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f636, plain, (spl30_45 | spl30_46 | spl30_47 | spl30_48), inference(avatar_split_clause, [], [f119, f633, f629, f625, f621])).
fof(f119, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f619, plain, (spl30_41 | spl30_42 | spl30_43 | spl30_44), inference(avatar_split_clause, [], [f120, f616, f612, f608, f604])).
fof(f120, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f602, plain, (spl30_37 | spl30_38 | spl30_39 | spl30_40), inference(avatar_split_clause, [], [f121, f599, f595, f591, f587])).
fof(f121, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f585, plain, (spl30_33 | spl30_34 | spl30_35 | spl30_36), inference(avatar_split_clause, [], [f122, f582, f578, f574, f570])).
fof(f122, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f568, plain, (spl30_29 | spl30_30 | spl30_31 | spl30_32), inference(avatar_split_clause, [], [f123, f565, f561, f557, f553])).
fof(f123, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f551, plain, (spl30_25 | spl30_26 | spl30_27 | spl30_28), inference(avatar_split_clause, [], [f124, f548, f544, f540, f536])).
fof(f124, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f534, plain, (spl30_21 | spl30_22 | spl30_23 | spl30_24), inference(avatar_split_clause, [], [f125, f531, f527, f523, f519])).
fof(f125, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f517, plain, (spl30_17 | spl30_18 | spl30_19 | spl30_20), inference(avatar_split_clause, [], [f126, f514, f510, f506, f502])).
fof(f126, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f500, plain, (spl30_13 | spl30_14 | spl30_15 | spl30_16), inference(avatar_split_clause, [], [f127, f497, f493, f489, f485])).
fof(f127, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f483, plain, (spl30_9 | spl30_10 | spl30_11 | spl30_12), inference(avatar_split_clause, [], [f128, f480, f476, f472, f468])).
fof(f128, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f466, plain, (spl30_5 | spl30_6 | spl30_7 | spl30_8), inference(avatar_split_clause, [], [f129, f463, f459, f455, f451])).
fof(f129, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f449, plain, (spl30_1 | spl30_2 | spl30_3 | spl30_4), inference(avatar_split_clause, [], [f130, f446, f442, f438, f434])).
fof(f130, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).