fof(f2884, plain, $false, inference(avatar_sat_refutation, [], [f541, f558, f609, f626, f694, f711, f728, f762, f780, f781, f784, f785, f787, f790, f791, f794, f795, f800, f801, f811, f845, f913, f930, f981, f998, f1015, f1032, f1049, f1084, f1085, f1089, f1091, f1093, f1095, f1096, f1098, f1104, f1115, f1120, f1124, f1131, f1138, f1153, f1157, f1161, f1168, f1172, f1180, f1181, f1182, f1183, f1184, f1189, f1190, f1195, f1200, f1207, f1208, f1209, f1212, f1222, f1226, f1230, f1237, f1241, f1249, f1250, f1251, f1252, f1253, f1254, f1256, f1257, f1259, f1495, f1516, f1534, f1962, f1989, f2010, f2019, f2033, f2034, f2041, f2054, f2064, f2074, f2079, f2089, f2095, f2098, f2104, f2115, f2124, f2138, f2140, f2141, f2143, f2144, f2153, f2167, f2182, f2185, f2224, f2225, f2226, f2227, f2232, f2233, f2244, f2249, f2250, f2251, f2254, f2272, f2282, f2288, f2291, f2307, f2309, f2315, f2325, f2327, f2349, f2357, f2363, f2364, f2365, f2366, f2371, f2372, f2374, f2379, f2384, f2389, f2398, f2434, f2476, f2478, f2496, f2498, f2500, f2502, f2514, f2524, f2554, f2571, f2575, f2579, f2584, f2608, f2616, f2630, f2637, f2657, f2658, f2659, f2670, f2686, f2697, f2727, f2737, f2738, f2747, f2752, f2756, f2771, f2792, f2804, f2808, f2820, f2825, f2837, f2842, f2853, f2856, f2862, f2867, f2873, f2883])).
fof(f2883, plain, (~ spl20_24 | ~ spl20_210 | ~ spl20_243 | spl20_247), inference(avatar_contradiction_clause, [], [f2882])).
fof(f2882, plain, ($false | (~ spl20_24 | ~ spl20_210 | ~ spl20_243 | spl20_247)), inference(subsumption_resolution, [], [f2881, f1912])).
fof(f1912, plain, ((e23 = h1(e13)) | ~ spl20_243), inference(avatar_component_clause, [], [f1911])).
fof(f1911, plain, (spl20_243 <=> (e23 = h1(e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_243])])).
fof(f2881, plain, (~ (e23 = h1(e13)) | (~ spl20_24 | ~ spl20_210 | spl20_247)), inference(forward_demodulation, [], [f2880, f1704])).
fof(f1704, plain, ((e23 = h3(e13)) | ~ spl20_210), inference(avatar_component_clause, [], [f1703])).
fof(f1703, plain, (spl20_210 <=> (e23 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_210])])).
fof(f2880, plain, (~ (h1(e13) = h3(e13)) | (~ spl20_24 | spl20_247)), inference(forward_demodulation, [], [f1929, f608])).
fof(f608, plain, ((e13 = op1(e12, e12)) | ~ spl20_24), inference(avatar_component_clause, [], [f606])).
fof(f606, plain, (spl20_24 <=> (e13 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_24])])).
fof(f1929, plain, (~ (h3(e13) = h1(op1(e12, e12))) | spl20_247), inference(avatar_component_clause, [], [f1927])).
fof(f1927, plain, (spl20_247 <=> (h3(e13) = h1(op1(e12, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl20_247])])).
fof(f2873, plain, (~ spl20_25 | ~ spl20_89 | ~ spl20_178 | spl20_248), inference(avatar_contradiction_clause, [], [f2872])).
fof(f2872, plain, ($false | (~ spl20_25 | ~ spl20_89 | ~ spl20_178 | spl20_248)), inference(subsumption_resolution, [], [f2871, f917])).
fof(f917, plain, ((e20 = op2(e22, e21)) | ~ spl20_89), inference(avatar_component_clause, [], [f915])).
fof(f915, plain, (spl20_89 <=> (e20 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_89])])).
fof(f2871, plain, (~ (e20 = op2(e22, e21)) | (~ spl20_25 | ~ spl20_178 | spl20_248)), inference(forward_demodulation, [], [f2870, f428])).
fof(f428, plain, (e20 = h1(e10)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ((op2(e20, e20) = h1(e13)) & (op2(op2(e20, e20), e20) = h1(e12)) & (op2(op2(op2(e20, e20), e20), e20) = h1(e11)) & (e20 = h1(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax14)).
fof(f2870, plain, (~ (op2(e22, e21) = h1(e10)) | (~ spl20_25 | ~ spl20_178 | spl20_248)), inference(forward_demodulation, [], [f2869, f613])).
fof(f613, plain, ((e10 = op1(e12, e11)) | ~ spl20_25), inference(avatar_component_clause, [], [f611])).
fof(f611, plain, (spl20_25 <=> (e10 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_25])])).
fof(f2869, plain, (~ (op2(e22, e21) = h1(op1(e12, e11))) | (~ spl20_178 | spl20_248)), inference(forward_demodulation, [], [f1933, f1514])).
fof(f1514, plain, ((e21 = h1(e11)) | ~ spl20_178), inference(avatar_component_clause, [], [f1513])).
fof(f1513, plain, (spl20_178 <=> (e21 = h1(e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_178])])).
fof(f1933, plain, (~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | spl20_248), inference(avatar_component_clause, [], [f1931])).
fof(f1931, plain, (spl20_248 <=> (h1(op1(e12, e11)) = op2(e22, h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl20_248])])).
fof(f2867, plain, (~ spl20_19 | ~ spl20_83 | ~ spl20_243 | spl20_246), inference(avatar_contradiction_clause, [], [f2866])).
fof(f2866, plain, ($false | (~ spl20_19 | ~ spl20_83 | ~ spl20_243 | spl20_246)), inference(subsumption_resolution, [], [f2865, f891])).
fof(f891, plain, ((e22 = op2(e22, e23)) | ~ spl20_83), inference(avatar_component_clause, [], [f889])).
fof(f889, plain, (spl20_83 <=> (e22 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_83])])).
fof(f2865, plain, (~ (e22 = op2(e22, e23)) | (~ spl20_19 | ~ spl20_243 | spl20_246)), inference(forward_demodulation, [], [f2864, f1268])).
fof(f1268, plain, (e22 = h1(e12)), inference(forward_demodulation, [], [f1267, f1266])).
fof(f1266, plain, (e22 = op2(h1(e13), e20)), inference(backward_demodulation, [], [f426, f431])).
fof(f431, plain, (op2(e20, e20) = h1(e13)), inference(cnf_transformation, [], [f14])).
fof(f426, plain, (e22 = op2(op2(e20, e20), e20)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((op2(e20, e20) = e23) & (e22 = op2(op2(e20, e20), e20)) & (e21 = op2(op2(op2(e20, e20), e20), e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax13)).
fof(f1267, plain, (h1(e12) = op2(h1(e13), e20)), inference(forward_demodulation, [], [f430, f431])).
fof(f430, plain, (op2(op2(e20, e20), e20) = h1(e12)), inference(cnf_transformation, [], [f14])).
fof(f2864, plain, (~ (op2(e22, e23) = h1(e12)) | (~ spl20_19 | ~ spl20_243 | spl20_246)), inference(forward_demodulation, [], [f2863, f587])).
fof(f587, plain, ((e12 = op1(e12, e13)) | ~ spl20_19), inference(avatar_component_clause, [], [f585])).
fof(f585, plain, (spl20_19 <=> (e12 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_19])])).
fof(f2863, plain, (~ (op2(e22, e23) = h1(op1(e12, e13))) | (~ spl20_243 | spl20_246)), inference(forward_demodulation, [], [f1925, f1912])).
fof(f1925, plain, (~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | spl20_246), inference(avatar_component_clause, [], [f1923])).
fof(f1923, plain, (spl20_246 <=> (h1(op1(e12, e13)) = op2(e22, h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl20_246])])).
fof(f2862, plain, (~ spl20_15 | spl20_245), inference(avatar_contradiction_clause, [], [f2861])).
fof(f2861, plain, ($false | (~ spl20_15 | spl20_245)), inference(subsumption_resolution, [], [f2860, f1268])).
fof(f2860, plain, (~ (e22 = h1(e12)) | (~ spl20_15 | spl20_245)), inference(forward_demodulation, [], [f1921, f570])).
fof(f570, plain, ((e12 = op1(e13, e10)) | ~ spl20_15), inference(avatar_component_clause, [], [f568])).
fof(f568, plain, (spl20_15 <=> (e12 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_15])])).
fof(f1921, plain, (~ (e22 = h1(op1(e13, e10))) | spl20_245), inference(avatar_component_clause, [], [f1919])).
fof(f1919, plain, (spl20_245 <=> (e22 = h1(op1(e13, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl20_245])])).
fof(f2856, plain, (~ spl20_30 | spl20_249), inference(avatar_contradiction_clause, [], [f2855])).
fof(f2855, plain, ($false | (~ spl20_30 | spl20_249)), inference(trivial_inequality_removal, [], [f2854])).
fof(f2854, plain, (~ (h1(e11) = h1(e11)) | (~ spl20_30 | spl20_249)), inference(forward_demodulation, [], [f1937, f634])).
fof(f634, plain, ((e11 = op1(e12, e10)) | ~ spl20_30), inference(avatar_component_clause, [], [f632])).
fof(f632, plain, (spl20_30 <=> (e11 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_30])])).
fof(f1937, plain, (~ (h1(e11) = h1(op1(e12, e10))) | spl20_249), inference(avatar_component_clause, [], [f1935])).
fof(f1935, plain, (spl20_249 <=> (h1(e11) = h1(op1(e12, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl20_249])])).
fof(f2853, plain, (~ spl20_5 | ~ spl20_69 | ~ spl20_243 | spl20_244), inference(avatar_contradiction_clause, [], [f2852])).
fof(f2852, plain, ($false | (~ spl20_5 | ~ spl20_69 | ~ spl20_243 | spl20_244)), inference(subsumption_resolution, [], [f2851, f832])).
fof(f832, plain, ((e20 = op2(e23, e22)) | ~ spl20_69), inference(avatar_component_clause, [], [f830])).
fof(f830, plain, (spl20_69 <=> (e20 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_69])])).
fof(f2851, plain, (~ (e20 = op2(e23, e22)) | (~ spl20_5 | ~ spl20_243 | spl20_244)), inference(forward_demodulation, [], [f2850, f428])).
fof(f2850, plain, (~ (op2(e23, e22) = h1(e10)) | (~ spl20_5 | ~ spl20_243 | spl20_244)), inference(forward_demodulation, [], [f2849, f528])).
fof(f528, plain, ((e10 = op1(e13, e12)) | ~ spl20_5), inference(avatar_component_clause, [], [f526])).
fof(f526, plain, (spl20_5 <=> (e10 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_5])])).
fof(f2849, plain, (~ (op2(e23, e22) = h1(op1(e13, e12))) | (~ spl20_243 | spl20_244)), inference(forward_demodulation, [], [f1917, f1912])).
fof(f1917, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | spl20_244), inference(avatar_component_clause, [], [f1915])).
fof(f1915, plain, (spl20_244 <=> (h1(op1(e13, e12)) = op2(h1(e13), e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_244])])).
fof(f2842, plain, (~ spl20_38 | ~ spl20_102 | ~ spl20_178 | spl20_250), inference(avatar_contradiction_clause, [], [f2841])).
fof(f2841, plain, ($false | (~ spl20_38 | ~ spl20_102 | ~ spl20_178 | spl20_250)), inference(subsumption_resolution, [], [f2840, f972])).
fof(f972, plain, ((e21 = op2(e21, e22)) | ~ spl20_102), inference(avatar_component_clause, [], [f970])).
fof(f970, plain, (spl20_102 <=> (e21 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_102])])).
fof(f2840, plain, (~ (e21 = op2(e21, e22)) | (~ spl20_38 | ~ spl20_178 | spl20_250)), inference(forward_demodulation, [], [f2839, f1514])).
fof(f2839, plain, (~ (op2(e21, e22) = h1(e11)) | (~ spl20_38 | ~ spl20_178 | spl20_250)), inference(forward_demodulation, [], [f2838, f668])).
fof(f668, plain, ((e11 = op1(e11, e12)) | ~ spl20_38), inference(avatar_component_clause, [], [f666])).
fof(f666, plain, (spl20_38 <=> (e11 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_38])])).
fof(f2838, plain, (~ (op2(e21, e22) = h1(op1(e11, e12))) | (~ spl20_178 | spl20_250)), inference(forward_demodulation, [], [f1941, f1514])).
fof(f1941, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | spl20_250), inference(avatar_component_clause, [], [f1939])).
fof(f1939, plain, (spl20_250 <=> (h1(op1(e11, e12)) = op2(h1(e11), e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_250])])).
fof(f2837, plain, (~ spl20_2 | ~ spl20_66 | ~ spl20_178 | spl20_242 | ~ spl20_243), inference(avatar_contradiction_clause, [], [f2836])).
fof(f2836, plain, ($false | (~ spl20_2 | ~ spl20_66 | ~ spl20_178 | spl20_242 | ~ spl20_243)), inference(subsumption_resolution, [], [f2835, f819])).
fof(f819, plain, ((e21 = op2(e23, e23)) | ~ spl20_66), inference(avatar_component_clause, [], [f817])).
fof(f817, plain, (spl20_66 <=> (e21 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_66])])).
fof(f2835, plain, (~ (e21 = op2(e23, e23)) | (~ spl20_2 | ~ spl20_178 | spl20_242 | ~ spl20_243)), inference(forward_demodulation, [], [f2834, f1514])).
fof(f2834, plain, (~ (op2(e23, e23) = h1(e11)) | (~ spl20_2 | spl20_242 | ~ spl20_243)), inference(forward_demodulation, [], [f2833, f515])).
fof(f515, plain, ((e11 = op1(e13, e13)) | ~ spl20_2), inference(avatar_component_clause, [], [f513])).
fof(f513, plain, (spl20_2 <=> (e11 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_2])])).
fof(f2833, plain, (~ (op2(e23, e23) = h1(op1(e13, e13))) | (spl20_242 | ~ spl20_243)), inference(forward_demodulation, [], [f1909, f1912])).
fof(f1909, plain, (~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | spl20_242), inference(avatar_component_clause, [], [f1907])).
fof(f1907, plain, (spl20_242 <=> (h1(op1(e13, e13)) = op2(h1(e13), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl20_242])])).
fof(f2825, plain, (~ spl20_49 | ~ spl20_113 | ~ spl20_243 | spl20_252), inference(avatar_contradiction_clause, [], [f2824])).
fof(f2824, plain, ($false | (~ spl20_49 | ~ spl20_113 | ~ spl20_243 | spl20_252)), inference(subsumption_resolution, [], [f2823, f1019])).
fof(f1019, plain, ((e20 = op2(e20, e23)) | ~ spl20_113), inference(avatar_component_clause, [], [f1017])).
fof(f1017, plain, (spl20_113 <=> (e20 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_113])])).
fof(f2823, plain, (~ (e20 = op2(e20, e23)) | (~ spl20_49 | ~ spl20_243 | spl20_252)), inference(forward_demodulation, [], [f2822, f428])).
fof(f2822, plain, (~ (op2(e20, e23) = h1(e10)) | (~ spl20_49 | ~ spl20_243 | spl20_252)), inference(forward_demodulation, [], [f2821, f715])).
fof(f715, plain, ((e10 = op1(e10, e13)) | ~ spl20_49), inference(avatar_component_clause, [], [f713])).
fof(f713, plain, (spl20_49 <=> (e10 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_49])])).
fof(f2821, plain, (~ (op2(e20, e23) = h1(op1(e10, e13))) | (~ spl20_243 | spl20_252)), inference(forward_demodulation, [], [f1949, f1912])).
fof(f1949, plain, (~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | spl20_252), inference(avatar_component_clause, [], [f1947])).
fof(f1947, plain, (spl20_252 <=> (h1(op1(e10, e13)) = op2(e20, h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl20_252])])).
fof(f2820, plain, (~ spl20_12 | ~ spl20_76 | ~ spl20_178 | spl20_241 | ~ spl20_243), inference(avatar_contradiction_clause, [], [f2819])).
fof(f2819, plain, ($false | (~ spl20_12 | ~ spl20_76 | ~ spl20_178 | spl20_241 | ~ spl20_243)), inference(subsumption_resolution, [], [f2818, f861])).
fof(f861, plain, ((e23 = op2(e23, e21)) | ~ spl20_76), inference(avatar_component_clause, [], [f859])).
fof(f859, plain, (spl20_76 <=> (e23 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_76])])).
fof(f2818, plain, (~ (e23 = op2(e23, e21)) | (~ spl20_12 | ~ spl20_178 | spl20_241 | ~ spl20_243)), inference(forward_demodulation, [], [f2817, f1912])).
fof(f2817, plain, (~ (op2(e23, e21) = h1(e13)) | (~ spl20_12 | ~ spl20_178 | spl20_241 | ~ spl20_243)), inference(forward_demodulation, [], [f2816, f557])).
fof(f557, plain, ((e13 = op1(e13, e11)) | ~ spl20_12), inference(avatar_component_clause, [], [f555])).
fof(f555, plain, (spl20_12 <=> (e13 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_12])])).
fof(f2816, plain, (~ (op2(e23, e21) = h1(op1(e13, e11))) | (~ spl20_178 | spl20_241 | ~ spl20_243)), inference(forward_demodulation, [], [f2815, f1912])).
fof(f2815, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), e21)) | (~ spl20_178 | spl20_241)), inference(forward_demodulation, [], [f1905, f1514])).
fof(f1905, plain, (~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | spl20_241), inference(avatar_component_clause, [], [f1903])).
fof(f1903, plain, (spl20_241 <=> (h1(op1(e13, e11)) = op2(h1(e13), h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl20_241])])).
fof(f2808, plain, (~ spl20_55 | ~ spl20_119 | spl20_253), inference(avatar_contradiction_clause, [], [f2807])).
fof(f2807, plain, ($false | (~ spl20_55 | ~ spl20_119 | spl20_253)), inference(subsumption_resolution, [], [f2806, f1044])).
fof(f1044, plain, ((e22 = op2(e20, e22)) | ~ spl20_119), inference(avatar_component_clause, [], [f1042])).
fof(f1042, plain, (spl20_119 <=> (e22 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_119])])).
fof(f2806, plain, (~ (e22 = op2(e20, e22)) | (~ spl20_55 | spl20_253)), inference(forward_demodulation, [], [f2805, f1268])).
fof(f2805, plain, (~ (op2(e20, e22) = h1(e12)) | (~ spl20_55 | spl20_253)), inference(forward_demodulation, [], [f1953, f740])).
fof(f740, plain, ((e12 = op1(e10, e12)) | ~ spl20_55), inference(avatar_component_clause, [], [f738])).
fof(f738, plain, (spl20_55 <=> (e12 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_55])])).
fof(f1953, plain, (~ (op2(e20, e22) = h1(op1(e10, e12))) | spl20_253), inference(avatar_component_clause, [], [f1951])).
fof(f1951, plain, (spl20_253 <=> (op2(e20, e22) = h1(op1(e10, e12)))), introduced(avatar_definition, [new_symbols(naming, [spl20_253])])).
fof(f2804, plain, (~ spl20_36 | ~ spl20_100 | ~ spl20_178 | spl20_240 | ~ spl20_243), inference(avatar_contradiction_clause, [], [f2803])).
fof(f2803, plain, ($false | (~ spl20_36 | ~ spl20_100 | ~ spl20_178 | spl20_240 | ~ spl20_243)), inference(subsumption_resolution, [], [f2802, f963])).
fof(f963, plain, ((e23 = op2(e21, e23)) | ~ spl20_100), inference(avatar_component_clause, [], [f961])).
fof(f961, plain, (spl20_100 <=> (e23 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_100])])).
fof(f2802, plain, (~ (e23 = op2(e21, e23)) | (~ spl20_36 | ~ spl20_178 | spl20_240 | ~ spl20_243)), inference(forward_demodulation, [], [f2801, f1912])).
fof(f2801, plain, (~ (op2(e21, e23) = h1(e13)) | (~ spl20_36 | ~ spl20_178 | spl20_240 | ~ spl20_243)), inference(forward_demodulation, [], [f2800, f659])).
fof(f659, plain, ((e13 = op1(e11, e13)) | ~ spl20_36), inference(avatar_component_clause, [], [f657])).
fof(f657, plain, (spl20_36 <=> (e13 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_36])])).
fof(f2800, plain, (~ (op2(e21, e23) = h1(op1(e11, e13))) | (~ spl20_178 | spl20_240 | ~ spl20_243)), inference(forward_demodulation, [], [f2799, f1514])).
fof(f2799, plain, (~ (h1(op1(e11, e13)) = op2(h1(e11), e23)) | (spl20_240 | ~ spl20_243)), inference(forward_demodulation, [], [f1901, f1912])).
fof(f1901, plain, (~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | spl20_240), inference(avatar_component_clause, [], [f1899])).
fof(f1899, plain, (spl20_240 <=> (h1(op1(e11, e13)) = op2(h1(e11), h1(e13)))), introduced(avatar_definition, [new_symbols(naming, [spl20_240])])).
fof(f2792, plain, (~ spl20_58 | ~ spl20_122 | ~ spl20_178 | spl20_254), inference(avatar_contradiction_clause, [], [f2791])).
fof(f2791, plain, ($false | (~ spl20_58 | ~ spl20_122 | ~ spl20_178 | spl20_254)), inference(subsumption_resolution, [], [f2790, f1057])).
fof(f1057, plain, ((e21 = op2(e20, e21)) | ~ spl20_122), inference(avatar_component_clause, [], [f1055])).
fof(f1055, plain, (spl20_122 <=> (e21 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_122])])).
fof(f2790, plain, (~ (e21 = op2(e20, e21)) | (~ spl20_58 | ~ spl20_178 | spl20_254)), inference(forward_demodulation, [], [f2789, f1514])).
fof(f2789, plain, (~ (op2(e20, e21) = h1(e11)) | (~ spl20_58 | ~ spl20_178 | spl20_254)), inference(forward_demodulation, [], [f2788, f753])).
fof(f753, plain, ((e11 = op1(e10, e11)) | ~ spl20_58), inference(avatar_component_clause, [], [f751])).
fof(f751, plain, (spl20_58 <=> (e11 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_58])])).
fof(f2788, plain, (~ (op2(e20, e21) = h1(op1(e10, e11))) | (~ spl20_178 | spl20_254)), inference(forward_demodulation, [], [f1957, f1514])).
fof(f1957, plain, (~ (h1(op1(e10, e11)) = op2(e20, h1(e11))) | spl20_254), inference(avatar_component_clause, [], [f1955])).
fof(f1955, plain, (spl20_254 <=> (h1(op1(e10, e11)) = op2(e20, h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl20_254])])).
fof(f2771, plain, (~ spl20_10 | ~ spl20_58), inference(avatar_split_clause, [], [f2770, f751, f547])).
fof(f547, plain, (spl20_10 <=> (e11 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_10])])).
fof(f2770, plain, (~ (e11 = op1(e13, e11)) | ~ spl20_58), inference(backward_demodulation, [], [f168, f753])).
fof(f168, plain, ~ (op1(e10, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f5, plain, (~ (op1(e13, e12) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e13)) & ~ (op1(e13, e11) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e13)) & ~ (op1(e13, e10) = op1(e13, e12)) & ~ (op1(e13, e10) = op1(e13, e11)) & ~ (op1(e12, e12) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e13)) & ~ (op1(e12, e11) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e13)) & ~ (op1(e12, e10) = op1(e12, e12)) & ~ (op1(e12, e10) = op1(e12, e11)) & ~ (op1(e11, e12) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e13)) & ~ (op1(e11, e11) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e13)) & ~ (op1(e11, e10) = op1(e11, e12)) & ~ (op1(e11, e10) = op1(e11, e11)) & ~ (op1(e10, e12) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e13)) & ~ (op1(e10, e11) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e13)) & ~ (op1(e10, e10) = op1(e10, e12)) & ~ (op1(e10, e10) = op1(e10, e11)) & ~ (op1(e12, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e13, e13)) & ~ (op1(e11, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e13, e13)) & ~ (op1(e10, e13) = op1(e12, e13)) & ~ (op1(e10, e13) = op1(e11, e13)) & ~ (op1(e12, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e13, e12)) & ~ (op1(e11, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e13, e12)) & ~ (op1(e10, e12) = op1(e12, e12)) & ~ (op1(e10, e12) = op1(e11, e12)) & ~ (op1(e12, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e13, e11)) & ~ (op1(e11, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e13, e11)) & ~ (op1(e10, e11) = op1(e12, e11)) & ~ (op1(e10, e11) = op1(e11, e11)) & ~ (op1(e12, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e13, e10)) & ~ (op1(e11, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e13, e10)) & ~ (op1(e10, e10) = op1(e12, e10)) & ~ (op1(e10, e10) = op1(e11, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax5)).
fof(f2756, plain, (~ spl20_64 | spl20_255), inference(avatar_contradiction_clause, [], [f2755])).
fof(f2755, plain, ($false | (~ spl20_64 | spl20_255)), inference(trivial_inequality_removal, [], [f2754])).
fof(f2754, plain, (~ (h1(e13) = h1(e13)) | (~ spl20_64 | spl20_255)), inference(forward_demodulation, [], [f1961, f778])).
fof(f778, plain, ((op1(e10, e10) = e13) | ~ spl20_64), inference(avatar_component_clause, [], [f776])).
fof(f776, plain, (spl20_64 <=> (op1(e10, e10) = e13)), introduced(avatar_definition, [new_symbols(naming, [spl20_64])])).
fof(f1961, plain, (~ (h1(e13) = h1(op1(e10, e10))) | spl20_255), inference(avatar_component_clause, [], [f1959])).
fof(f1959, plain, (spl20_255 <=> (h1(e13) = h1(op1(e10, e10)))), introduced(avatar_definition, [new_symbols(naming, [spl20_255])])).
fof(f2752, plain, (~ spl20_5 | ~ spl20_8), inference(avatar_contradiction_clause, [], [f2751])).
fof(f2751, plain, ($false | (~ spl20_5 | ~ spl20_8)), inference(subsumption_resolution, [], [f2750, f258])).
fof(f258, plain, ~ (e10 = e13), inference(cnf_transformation, [], [f7])).
fof(f7, plain, (~ (e12 = e13) & ~ (e11 = e13) & ~ (e11 = e12) & ~ (e10 = e13) & ~ (e10 = e12) & ~ (e10 = e11)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax7)).
fof(f2750, plain, ((e10 = e13) | (~ spl20_5 | ~ spl20_8)), inference(forward_demodulation, [], [f540, f528])).
fof(f540, plain, ((e13 = op1(e13, e12)) | ~ spl20_8), inference(avatar_component_clause, [], [f538])).
fof(f538, plain, (spl20_8 <=> (e13 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_8])])).
fof(f2747, plain, (~ spl20_30 | ~ spl20_31), inference(avatar_contradiction_clause, [], [f2746])).
fof(f2746, plain, ($false | (~ spl20_30 | ~ spl20_31)), inference(subsumption_resolution, [], [f2745, f259])).
fof(f259, plain, ~ (e11 = e12), inference(cnf_transformation, [], [f7])).
fof(f2745, plain, ((e11 = e12) | (~ spl20_30 | ~ spl20_31)), inference(forward_demodulation, [], [f638, f634])).
fof(f638, plain, ((e12 = op1(e12, e10)) | ~ spl20_31), inference(avatar_component_clause, [], [f636])).
fof(f636, plain, (spl20_31 <=> (e12 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_31])])).
fof(f2738, plain, (~ spl20_36 | ~ spl20_44), inference(avatar_split_clause, [], [f2732, f691, f657])).
fof(f691, plain, (spl20_44 <=> (e13 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_44])])).
fof(f2732, plain, (~ (e13 = op1(e11, e13)) | ~ spl20_44), inference(backward_demodulation, [], [f194, f693])).
fof(f693, plain, ((e13 = op1(e11, e11)) | ~ spl20_44), inference(avatar_component_clause, [], [f691])).
fof(f194, plain, ~ (op1(e11, e11) = op1(e11, e13)), inference(cnf_transformation, [], [f5])).
fof(f2737, plain, (~ spl20_12 | ~ spl20_44), inference(avatar_split_clause, [], [f2731, f691, f555])).
fof(f2731, plain, (~ (e13 = op1(e13, e11)) | ~ spl20_44), inference(backward_demodulation, [], [f170, f693])).
fof(f170, plain, ~ (op1(e11, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f2727, plain, (~ spl20_43 | ~ spl20_59), inference(avatar_split_clause, [], [f2725, f755, f687])).
fof(f687, plain, (spl20_43 <=> (e12 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_43])])).
fof(f755, plain, (spl20_59 <=> (e12 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_59])])).
fof(f2725, plain, (~ (e12 = op1(e11, e11)) | ~ spl20_59), inference(backward_demodulation, [], [f166, f757])).
fof(f757, plain, ((e12 = op1(e10, e11)) | ~ spl20_59), inference(avatar_component_clause, [], [f755])).
fof(f166, plain, ~ (op1(e10, e11) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f2697, plain, (~ spl20_20 | ~ spl20_24), inference(avatar_split_clause, [], [f2696, f606, f589])).
fof(f589, plain, (spl20_20 <=> (e13 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_20])])).
fof(f2696, plain, (~ (e13 = op1(e12, e13)) | ~ spl20_24), inference(forward_demodulation, [], [f201, f608])).
fof(f201, plain, ~ (op1(e12, e12) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f2686, plain, (~ spl20_89 | ~ spl20_91), inference(avatar_contradiction_clause, [], [f2685])).
fof(f2685, plain, ($false | (~ spl20_89 | ~ spl20_91)), inference(subsumption_resolution, [], [f2684, f263])).
fof(f263, plain, ~ (e20 = e22), inference(cnf_transformation, [], [f8])).
fof(f8, plain, (~ (e22 = e23) & ~ (e21 = e23) & ~ (e21 = e22) & ~ (e20 = e23) & ~ (e20 = e22) & ~ (e20 = e21)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax8)).
fof(f2684, plain, ((e20 = e22) | (~ spl20_89 | ~ spl20_91)), inference(backward_demodulation, [], [f925, f917])).
fof(f925, plain, ((e22 = op2(e22, e21)) | ~ spl20_91), inference(avatar_component_clause, [], [f923])).
fof(f923, plain, (spl20_91 <=> (e22 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_91])])).
fof(f2670, plain, (~ spl20_99 | ~ spl20_162), inference(avatar_split_clause, [], [f2666, f1429, f957])).
fof(f957, plain, (spl20_99 <=> (e22 = op2(e21, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_99])])).
fof(f1429, plain, (spl20_162 <=> (e22 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_162])])).
fof(f2666, plain, (~ (e22 = op2(e21, e23)) | ~ spl20_162), inference(backward_demodulation, [], [f1282, f1430])).
fof(f1430, plain, ((e22 = h2(e13)) | ~ spl20_162), inference(avatar_component_clause, [], [f1429])).
fof(f1282, plain, ~ (op2(e21, e23) = h2(e13)), inference(backward_demodulation, [], [f242, f435])).
fof(f435, plain, (op2(e21, e21) = h2(e13)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((op2(e21, e21) = h2(e13)) & (op2(op2(e21, e21), e21) = h2(e12)) & (h2(e11) = op2(op2(op2(e21, e21), e21), e21)) & (e21 = h2(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax15)).
fof(f242, plain, ~ (op2(e21, e21) = op2(e21, e23)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, (~ (op2(e23, e22) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e23)) & ~ (op2(e23, e21) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e23)) & ~ (op2(e23, e20) = op2(e23, e22)) & ~ (op2(e23, e20) = op2(e23, e21)) & ~ (op2(e22, e22) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e23)) & ~ (op2(e22, e21) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e23)) & ~ (op2(e22, e20) = op2(e22, e22)) & ~ (op2(e22, e20) = op2(e22, e21)) & ~ (op2(e21, e22) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e23)) & ~ (op2(e21, e21) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e23)) & ~ (op2(e21, e20) = op2(e21, e22)) & ~ (op2(e21, e20) = op2(e21, e21)) & ~ (op2(e20, e22) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e23)) & ~ (op2(e20, e21) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e23)) & ~ (op2(e20, e20) = op2(e20, e22)) & ~ (op2(e20, e20) = op2(e20, e21)) & ~ (op2(e22, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e23, e23)) & ~ (op2(e21, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e23, e23)) & ~ (op2(e20, e23) = op2(e22, e23)) & ~ (op2(e20, e23) = op2(e21, e23)) & ~ (op2(e22, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e23, e22)) & ~ (op2(e21, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e23, e22)) & ~ (op2(e20, e22) = op2(e22, e22)) & ~ (op2(e20, e22) = op2(e21, e22)) & ~ (op2(e22, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e23, e21)) & ~ (op2(e21, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e23, e21)) & ~ (op2(e20, e21) = op2(e22, e21)) & ~ (op2(e20, e21) = op2(e21, e21)) & ~ (op2(e22, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e23, e20)) & ~ (op2(e21, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e23, e20)) & ~ (op2(e20, e20) = op2(e22, e20)) & ~ (op2(e20, e20) = op2(e21, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax6)).
fof(f2659, plain, (spl20_210 | ~ spl20_88), inference(avatar_split_clause, [], [f2617, f910, f1703])).
fof(f910, plain, (spl20_88 <=> (e23 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_88])])).
fof(f2617, plain, ((e23 = h3(e13)) | ~ spl20_88), inference(backward_demodulation, [], [f439, f912])).
fof(f912, plain, ((e23 = op2(e22, e22)) | ~ spl20_88), inference(avatar_component_clause, [], [f910])).
fof(f439, plain, (op2(e22, e22) = h3(e13)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ((op2(e22, e22) = h3(e13)) & (op2(op2(e22, e22), e22) = h3(e12)) & (h3(e11) = op2(op2(op2(e22, e22), e22), e22)) & (e22 = h3(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax16)).
fof(f2658, plain, (~ spl20_162 | ~ spl20_91), inference(avatar_split_clause, [], [f2587, f923, f1429])).
fof(f2587, plain, (~ (e22 = h2(e13)) | ~ spl20_91), inference(forward_demodulation, [], [f1278, f925])).
fof(f1278, plain, ~ (op2(e22, e21) = h2(e13)), inference(backward_demodulation, [], [f217, f435])).
fof(f217, plain, ~ (op2(e21, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f2657, plain, (~ spl20_162 | ~ spl20_43 | ~ spl20_178 | spl20_239), inference(avatar_split_clause, [], [f2639, f1895, f1513, f687, f1429])).
fof(f1895, plain, (spl20_239 <=> (h1(op1(e11, e11)) = op2(h1(e11), h1(e11)))), introduced(avatar_definition, [new_symbols(naming, [spl20_239])])).
fof(f2639, plain, (~ (e22 = h2(e13)) | (~ spl20_43 | ~ spl20_178 | spl20_239)), inference(forward_demodulation, [], [f2636, f1268])).
fof(f2636, plain, (~ (h1(e12) = h2(e13)) | (~ spl20_43 | ~ spl20_178 | spl20_239)), inference(backward_demodulation, [], [f2590, f689])).
fof(f689, plain, ((e12 = op1(e11, e11)) | ~ spl20_43), inference(avatar_component_clause, [], [f687])).
fof(f2590, plain, (~ (h2(e13) = h1(op1(e11, e11))) | (~ spl20_178 | spl20_239)), inference(forward_demodulation, [], [f2589, f435])).
fof(f2589, plain, (~ (op2(e21, e21) = h1(op1(e11, e11))) | (~ spl20_178 | spl20_239)), inference(forward_demodulation, [], [f1897, f1514])).
fof(f1897, plain, (~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | spl20_239), inference(avatar_component_clause, [], [f1895])).
fof(f2637, plain, (~ spl20_27 | ~ spl20_43), inference(avatar_split_clause, [], [f2631, f687, f619])).
fof(f619, plain, (spl20_27 <=> (e12 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_27])])).
fof(f2631, plain, (~ (e12 = op1(e12, e11)) | ~ spl20_43), inference(backward_demodulation, [], [f169, f689])).
fof(f169, plain, ~ (op1(e11, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f2630, plain, (~ spl20_41 | ~ spl20_45), inference(avatar_split_clause, [], [f2628, f696, f679])).
fof(f679, plain, (spl20_41 <=> (e10 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_41])])).
fof(f696, plain, (spl20_45 <=> (e10 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_45])])).
fof(f2628, plain, (~ (e10 = op1(e11, e11)) | ~ spl20_45), inference(backward_demodulation, [], [f190, f698])).
fof(f698, plain, ((e10 = op1(e11, e10)) | ~ spl20_45), inference(avatar_component_clause, [], [f696])).
fof(f190, plain, ~ (op1(e11, e10) = op1(e11, e11)), inference(cnf_transformation, [], [f5])).
fof(f2616, plain, (~ spl20_94 | ~ spl20_96), inference(avatar_contradiction_clause, [], [f2615])).
fof(f2615, plain, ($false | (~ spl20_94 | ~ spl20_96)), inference(subsumption_resolution, [], [f2614, f266])).
fof(f266, plain, ~ (e21 = e23), inference(cnf_transformation, [], [f8])).
fof(f2614, plain, ((e21 = e23) | (~ spl20_94 | ~ spl20_96)), inference(forward_demodulation, [], [f946, f938])).
fof(f938, plain, ((e21 = op2(e22, e20)) | ~ spl20_94), inference(avatar_component_clause, [], [f936])).
fof(f936, plain, (spl20_94 <=> (e21 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_94])])).
fof(f946, plain, ((e23 = op2(e22, e20)) | ~ spl20_96), inference(avatar_component_clause, [], [f944])).
fof(f944, plain, (spl20_96 <=> (e23 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_96])])).
fof(f2608, plain, (~ spl20_108 | spl20_229), inference(avatar_contradiction_clause, [], [f2607])).
fof(f2607, plain, ($false | (~ spl20_108 | spl20_229)), inference(subsumption_resolution, [], [f2606, f1815])).
fof(f1815, plain, (~ (e23 = h2(e13)) | spl20_229), inference(avatar_component_clause, [], [f1813])).
fof(f1813, plain, (spl20_229 <=> (e23 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_229])])).
fof(f2606, plain, ((e23 = h2(e13)) | ~ spl20_108), inference(backward_demodulation, [], [f435, f997])).
fof(f997, plain, ((e23 = op2(e21, e21)) | ~ spl20_108), inference(avatar_component_clause, [], [f995])).
fof(f995, plain, (spl20_108 <=> (e23 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_108])])).
fof(f2584, plain, (~ spl20_116 | ~ spl20_243), inference(avatar_split_clause, [], [f2583, f1911, f1029])).
fof(f1029, plain, (spl20_116 <=> (e23 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_116])])).
fof(f2583, plain, (~ (e23 = op2(e20, e23)) | ~ spl20_243), inference(forward_demodulation, [], [f1265, f1912])).
fof(f1265, plain, ~ (op2(e20, e23) = h1(e13)), inference(backward_demodulation, [], [f234, f431])).
fof(f234, plain, ~ (op2(e20, e20) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f2579, plain, (~ spl20_111 | ~ spl20_79), inference(avatar_split_clause, [], [f2578, f872, f1008])).
fof(f1008, plain, (spl20_111 <=> (e22 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_111])])).
fof(f872, plain, (spl20_79 <=> (e22 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_79])])).
fof(f2578, plain, (~ (e22 = op2(e21, e20)) | ~ spl20_79), inference(forward_demodulation, [], [f212, f874])).
fof(f874, plain, ((e22 = op2(e23, e20)) | ~ spl20_79), inference(avatar_component_clause, [], [f872])).
fof(f212, plain, ~ (op2(e21, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f2575, plain, (~ spl20_13 | ~ spl20_15), inference(avatar_contradiction_clause, [], [f2574])).
fof(f2574, plain, ($false | (~ spl20_13 | ~ spl20_15)), inference(subsumption_resolution, [], [f2573, f257])).
fof(f257, plain, ~ (e10 = e12), inference(cnf_transformation, [], [f7])).
fof(f2573, plain, ((e10 = e12) | (~ spl20_13 | ~ spl20_15)), inference(backward_demodulation, [], [f570, f562])).
fof(f562, plain, ((e10 = op1(e13, e10)) | ~ spl20_13), inference(avatar_component_clause, [], [f560])).
fof(f560, plain, (spl20_13 <=> (e10 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_13])])).
fof(f2571, plain, (~ spl20_17 | ~ spl20_19), inference(avatar_contradiction_clause, [], [f2570])).
fof(f2570, plain, ($false | (~ spl20_17 | ~ spl20_19)), inference(subsumption_resolution, [], [f2569, f257])).
fof(f2569, plain, ((e10 = e12) | (~ spl20_17 | ~ spl20_19)), inference(backward_demodulation, [], [f587, f579])).
fof(f579, plain, ((e10 = op1(e12, e13)) | ~ spl20_17), inference(avatar_component_clause, [], [f577])).
fof(f577, plain, (spl20_17 <=> (e10 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_17])])).
fof(f2554, plain, (~ spl20_69 | ~ spl20_70), inference(avatar_contradiction_clause, [], [f2553])).
fof(f2553, plain, ($false | (~ spl20_69 | ~ spl20_70)), inference(subsumption_resolution, [], [f2552, f262])).
fof(f262, plain, ~ (e20 = e21), inference(cnf_transformation, [], [f8])).
fof(f2552, plain, ((e20 = e21) | (~ spl20_69 | ~ spl20_70)), inference(forward_demodulation, [], [f836, f832])).
fof(f836, plain, ((e21 = op2(e23, e22)) | ~ spl20_70), inference(avatar_component_clause, [], [f834])).
fof(f834, plain, (spl20_70 <=> (e21 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_70])])).
fof(f2524, plain, (~ spl20_89 | ~ spl20_158), inference(avatar_split_clause, [], [f2520, f1409, f915])).
fof(f1409, plain, (spl20_158 <=> (e20 = h3(e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_158])])).
fof(f2520, plain, (~ (e20 = op2(e22, e21)) | ~ spl20_158), inference(backward_demodulation, [], [f1289, f1410])).
fof(f1410, plain, ((e20 = h3(e13)) | ~ spl20_158), inference(avatar_component_clause, [], [f1409])).
fof(f1289, plain, ~ (op2(e22, e21) = h3(e13)), inference(backward_demodulation, [], [f247, f439])).
fof(f247, plain, ~ (op2(e22, e21) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f2514, plain, (~ spl20_109 | ~ spl20_170), inference(avatar_split_clause, [], [f2509, f1470, f1000])).
fof(f1000, plain, (spl20_109 <=> (e20 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_109])])).
fof(f1470, plain, (spl20_170 <=> (e20 = h2(e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_170])])).
fof(f2509, plain, (~ (e20 = op2(e21, e20)) | ~ spl20_170), inference(backward_demodulation, [], [f1280, f1471])).
fof(f1471, plain, ((e20 = h2(e13)) | ~ spl20_170), inference(avatar_component_clause, [], [f1470])).
fof(f1280, plain, ~ (op2(e21, e20) = h2(e13)), inference(backward_demodulation, [], [f238, f435])).
fof(f238, plain, ~ (op2(e21, e20) = op2(e21, e21)), inference(cnf_transformation, [], [f6])).
fof(f2502, plain, (~ spl20_112 | ~ spl20_243), inference(avatar_split_clause, [], [f2501, f1911, f1012])).
fof(f1012, plain, (spl20_112 <=> (e23 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_112])])).
fof(f2501, plain, (~ (e23 = op2(e21, e20)) | ~ spl20_243), inference(forward_demodulation, [], [f1260, f1912])).
fof(f1260, plain, ~ (op2(e21, e20) = h1(e13)), inference(backward_demodulation, [], [f208, f431])).
fof(f208, plain, ~ (op2(e20, e20) = op2(e21, e20)), inference(cnf_transformation, [], [f6])).
fof(f2500, plain, (~ spl20_110 | ~ spl20_178), inference(avatar_split_clause, [], [f2499, f1513, f1004])).
fof(f1004, plain, (spl20_110 <=> (e21 = op2(e21, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_110])])).
fof(f2499, plain, (~ (e21 = op2(e21, e20)) | ~ spl20_178), inference(forward_demodulation, [], [f1272, f1514])).
fof(f1272, plain, ~ (op2(e21, e20) = h1(e11)), inference(backward_demodulation, [], [f211, f1270])).
fof(f1270, plain, (op2(e22, e20) = h1(e11)), inference(forward_demodulation, [], [f1269, f1266])).
fof(f1269, plain, (h1(e11) = op2(op2(h1(e13), e20), e20)), inference(forward_demodulation, [], [f429, f431])).
fof(f429, plain, (op2(op2(op2(e20, e20), e20), e20) = h1(e11)), inference(cnf_transformation, [], [f14])).
fof(f211, plain, ~ (op2(e21, e20) = op2(e22, e20)), inference(cnf_transformation, [], [f6])).
fof(f2498, plain, (~ spl20_90 | ~ spl20_178), inference(avatar_split_clause, [], [f2497, f1513, f919])).
fof(f919, plain, (spl20_90 <=> (e21 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_90])])).
fof(f2497, plain, (~ (e21 = op2(e22, e21)) | ~ spl20_178), inference(forward_demodulation, [], [f1274, f1514])).
fof(f1274, plain, ~ (op2(e22, e21) = h1(e11)), inference(backward_demodulation, [], [f244, f1270])).
fof(f244, plain, ~ (op2(e22, e20) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f2496, plain, (~ spl20_92 | ~ spl20_76), inference(avatar_split_clause, [], [f2495, f859, f927])).
fof(f927, plain, (spl20_92 <=> (e23 = op2(e22, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_92])])).
fof(f2495, plain, (~ (e23 = op2(e22, e21)) | ~ spl20_76), inference(forward_demodulation, [], [f219, f861])).
fof(f219, plain, ~ (op2(e22, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f2478, plain, (~ spl20_46 | ~ spl20_30), inference(avatar_split_clause, [], [f2477, f632, f700])).
fof(f700, plain, (spl20_46 <=> (e11 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_46])])).
fof(f2477, plain, (~ (e11 = op1(e11, e10)) | ~ spl20_30), inference(forward_demodulation, [], [f163, f634])).
fof(f163, plain, ~ (op1(e11, e10) = op1(e12, e10)), inference(cnf_transformation, [], [f5])).
fof(f2476, plain, (~ spl20_47 | ~ spl20_15), inference(avatar_split_clause, [], [f2475, f568, f704])).
fof(f704, plain, (spl20_47 <=> (e12 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_47])])).
fof(f2475, plain, (~ (e12 = op1(e11, e10)) | ~ spl20_15), inference(forward_demodulation, [], [f164, f570])).
fof(f164, plain, ~ (op1(e11, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f2434, plain, (~ spl20_38 | ~ spl20_39), inference(avatar_contradiction_clause, [], [f2433])).
fof(f2433, plain, ($false | (~ spl20_38 | ~ spl20_39)), inference(subsumption_resolution, [], [f2432, f259])).
fof(f2432, plain, ((e11 = e12) | (~ spl20_38 | ~ spl20_39)), inference(backward_demodulation, [], [f672, f668])).
fof(f672, plain, ((e12 = op1(e11, e12)) | ~ spl20_39), inference(avatar_component_clause, [], [f670])).
fof(f670, plain, (spl20_39 <=> (e12 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_39])])).
fof(f2398, plain, (~ spl20_94 | ~ spl20_95), inference(avatar_contradiction_clause, [], [f2397])).
fof(f2397, plain, ($false | (~ spl20_94 | ~ spl20_95)), inference(subsumption_resolution, [], [f2396, f265])).
fof(f265, plain, ~ (e21 = e22), inference(cnf_transformation, [], [f8])).
fof(f2396, plain, ((e21 = e22) | (~ spl20_94 | ~ spl20_95)), inference(forward_demodulation, [], [f942, f938])).
fof(f942, plain, ((e22 = op2(e22, e20)) | ~ spl20_95), inference(avatar_component_clause, [], [f940])).
fof(f940, plain, (spl20_95 <=> (e22 = op2(e22, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_95])])).
fof(f2389, plain, (spl20_162 | ~ spl20_107), inference(avatar_split_clause, [], [f2388, f991, f1429])).
fof(f991, plain, (spl20_107 <=> (e22 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_107])])).
fof(f2388, plain, ((e22 = h2(e13)) | ~ spl20_107), inference(backward_demodulation, [], [f435, f993])).
fof(f993, plain, ((e22 = op2(e21, e21)) | ~ spl20_107), inference(avatar_component_clause, [], [f991])).
fof(f2384, plain, (~ spl20_103 | ~ spl20_119), inference(avatar_split_clause, [], [f2382, f1042, f974])).
fof(f974, plain, (spl20_103 <=> (e22 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_103])])).
fof(f2382, plain, (~ (e22 = op2(e21, e22)) | ~ spl20_119), inference(backward_demodulation, [], [f220, f1044])).
fof(f220, plain, ~ (op2(e20, e22) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f2379, plain, (~ spl20_118 | ~ spl20_122), inference(avatar_split_clause, [], [f2377, f1055, f1038])).
fof(f1038, plain, (spl20_118 <=> (e21 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_118])])).
fof(f2377, plain, (~ (e21 = op2(e20, e22)) | ~ spl20_122), inference(backward_demodulation, [], [f235, f1057])).
fof(f235, plain, ~ (op2(e20, e21) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f2374, plain, (~ spl20_45 | ~ spl20_109 | ~ spl20_178 | spl20_251), inference(avatar_contradiction_clause, [], [f2373])).
fof(f2373, plain, ($false | (~ spl20_45 | ~ spl20_109 | ~ spl20_178 | spl20_251)), inference(subsumption_resolution, [], [f2370, f1002])).
fof(f1002, plain, ((e20 = op2(e21, e20)) | ~ spl20_109), inference(avatar_component_clause, [], [f1000])).
fof(f2370, plain, (~ (e20 = op2(e21, e20)) | (~ spl20_45 | ~ spl20_178 | spl20_251)), inference(backward_demodulation, [], [f2187, f1514])).
fof(f2187, plain, (~ (e20 = op2(h1(e11), e20)) | (~ spl20_45 | spl20_251)), inference(forward_demodulation, [], [f2181, f428])).
fof(f2181, plain, (~ (h1(e10) = op2(h1(e11), e20)) | (~ spl20_45 | spl20_251)), inference(backward_demodulation, [], [f1945, f698])).
fof(f1945, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | spl20_251), inference(avatar_component_clause, [], [f1943])).
fof(f1943, plain, (spl20_251 <=> (h1(op1(e11, e10)) = op2(h1(e11), e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_251])])).
fof(f2372, plain, (~ spl20_86 | ~ spl20_178), inference(avatar_split_clause, [], [f2368, f1513, f902])).
fof(f902, plain, (spl20_86 <=> (e21 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_86])])).
fof(f2368, plain, (~ (e21 = op2(e22, e22)) | ~ spl20_178), inference(backward_demodulation, [], [f1275, f1514])).
fof(f1275, plain, ~ (op2(e22, e22) = h1(e11)), inference(backward_demodulation, [], [f245, f1270])).
fof(f245, plain, ~ (op2(e22, e20) = op2(e22, e22)), inference(cnf_transformation, [], [f6])).
fof(f2371, plain, (~ spl20_78 | ~ spl20_178), inference(avatar_split_clause, [], [f2367, f1513, f868])).
fof(f868, plain, (spl20_78 <=> (e21 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_78])])).
fof(f2367, plain, (~ (e21 = op2(e23, e20)) | ~ spl20_178), inference(backward_demodulation, [], [f1273, f1514])).
fof(f1273, plain, ~ (op2(e23, e20) = h1(e11)), inference(backward_demodulation, [], [f213, f1270])).
fof(f213, plain, ~ (op2(e22, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f2366, plain, (spl20_79 | ~ spl20_243), inference(avatar_split_clause, [], [f2361, f1911, f872])).
fof(f2361, plain, ((e22 = op2(e23, e20)) | ~ spl20_243), inference(backward_demodulation, [], [f1266, f1912])).
fof(f2365, plain, (~ spl20_120 | ~ spl20_243), inference(avatar_split_clause, [], [f2360, f1911, f1046])).
fof(f1046, plain, (spl20_120 <=> (e23 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_120])])).
fof(f2360, plain, (~ (e23 = op2(e20, e22)) | ~ spl20_243), inference(backward_demodulation, [], [f1264, f1912])).
fof(f1264, plain, ~ (op2(e20, e22) = h1(e13)), inference(backward_demodulation, [], [f233, f431])).
fof(f233, plain, ~ (op2(e20, e20) = op2(e20, e22)), inference(cnf_transformation, [], [f6])).
fof(f2364, plain, (~ spl20_124 | ~ spl20_243), inference(avatar_split_clause, [], [f2359, f1911, f1063])).
fof(f1063, plain, (spl20_124 <=> (e23 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_124])])).
fof(f2359, plain, (~ (e23 = op2(e20, e21)) | ~ spl20_243), inference(backward_demodulation, [], [f1263, f1912])).
fof(f1263, plain, ~ (op2(e20, e21) = h1(e13)), inference(backward_demodulation, [], [f232, f431])).
fof(f232, plain, ~ (op2(e20, e20) = op2(e20, e21)), inference(cnf_transformation, [], [f6])).
fof(f2363, plain, (~ spl20_80 | ~ spl20_243), inference(avatar_split_clause, [], [f2358, f1911, f876])).
fof(f876, plain, (spl20_80 <=> (e23 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_80])])).
fof(f2358, plain, (~ (e23 = op2(e23, e20)) | ~ spl20_243), inference(backward_demodulation, [], [f1262, f1912])).
fof(f1262, plain, ~ (op2(e23, e20) = h1(e13)), inference(backward_demodulation, [], [f210, f431])).
fof(f210, plain, ~ (op2(e20, e20) = op2(e23, e20)), inference(cnf_transformation, [], [f6])).
fof(f2357, plain, (~ spl20_121 | ~ spl20_89), inference(avatar_split_clause, [], [f2356, f915, f1051])).
fof(f1051, plain, (spl20_121 <=> (e20 = op2(e20, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_121])])).
fof(f2356, plain, (~ (e20 = op2(e20, e21)) | ~ spl20_89), inference(forward_demodulation, [], [f215, f917])).
fof(f215, plain, ~ (op2(e20, e21) = op2(e22, e21)), inference(cnf_transformation, [], [f6])).
fof(f2349, plain, (~ spl20_117 | ~ spl20_113), inference(avatar_split_clause, [], [f2348, f1017, f1034])).
fof(f1034, plain, (spl20_117 <=> (e20 = op2(e20, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_117])])).
fof(f2348, plain, (~ (e20 = op2(e20, e22)) | ~ spl20_113), inference(forward_demodulation, [], [f237, f1019])).
fof(f237, plain, ~ (op2(e20, e22) = op2(e20, e23)), inference(cnf_transformation, [], [f6])).
fof(f2327, plain, (~ spl20_79 | ~ spl20_71), inference(avatar_split_clause, [], [f2326, f838, f872])).
fof(f838, plain, (spl20_71 <=> (e22 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_71])])).
fof(f2326, plain, (~ (e22 = op2(e23, e20)) | ~ spl20_71), inference(forward_demodulation, [], [f251, f840])).
fof(f840, plain, ((e22 = op2(e23, e22)) | ~ spl20_71), inference(avatar_component_clause, [], [f838])).
fof(f251, plain, ~ (op2(e23, e20) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f2325, plain, (~ spl20_74 | ~ spl20_76), inference(avatar_contradiction_clause, [], [f2324])).
fof(f2324, plain, ($false | (~ spl20_74 | ~ spl20_76)), inference(subsumption_resolution, [], [f2323, f266])).
fof(f2323, plain, ((e21 = e23) | (~ spl20_74 | ~ spl20_76)), inference(forward_demodulation, [], [f861, f853])).
fof(f853, plain, ((e21 = op2(e23, e21)) | ~ spl20_74), inference(avatar_component_clause, [], [f851])).
fof(f851, plain, (spl20_74 <=> (e21 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_74])])).
fof(f2315, plain, (~ spl20_57 | ~ spl20_25), inference(avatar_split_clause, [], [f2314, f611, f747])).
fof(f747, plain, (spl20_57 <=> (e10 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_57])])).
fof(f2314, plain, (~ (e10 = op1(e10, e11)) | ~ spl20_25), inference(forward_demodulation, [], [f167, f613])).
fof(f167, plain, ~ (op1(e10, e11) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f2309, plain, (~ spl20_53 | ~ spl20_5), inference(avatar_split_clause, [], [f2308, f526, f730])).
fof(f730, plain, (spl20_53 <=> (e10 = op1(e10, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_53])])).
fof(f2308, plain, (~ (e10 = op1(e10, e12)) | ~ spl20_5), inference(forward_demodulation, [], [f174, f528])).
fof(f174, plain, ~ (op1(e10, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f2307, plain, (~ spl20_53 | ~ spl20_49), inference(avatar_split_clause, [], [f2306, f713, f730])).
fof(f2306, plain, (~ (e10 = op1(e10, e12)) | ~ spl20_49), inference(forward_demodulation, [], [f189, f715])).
fof(f189, plain, ~ (op1(e10, e12) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f2291, plain, (~ spl20_34 | ~ spl20_36), inference(avatar_contradiction_clause, [], [f2290])).
fof(f2290, plain, ($false | (~ spl20_34 | ~ spl20_36)), inference(subsumption_resolution, [], [f2289, f260])).
fof(f260, plain, ~ (e11 = e13), inference(cnf_transformation, [], [f7])).
fof(f2289, plain, ((e11 = e13) | (~ spl20_34 | ~ spl20_36)), inference(forward_demodulation, [], [f659, f651])).
fof(f651, plain, ((e11 = op1(e11, e13)) | ~ spl20_34), inference(avatar_component_clause, [], [f649])).
fof(f649, plain, (spl20_34 <=> (e11 = op1(e11, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_34])])).
fof(f2288, plain, (~ spl20_21 | ~ spl20_5), inference(avatar_split_clause, [], [f2287, f526, f594])).
fof(f594, plain, (spl20_21 <=> (e10 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_21])])).
fof(f2287, plain, (~ (e10 = op1(e12, e12)) | ~ spl20_5), inference(forward_demodulation, [], [f177, f528])).
fof(f177, plain, ~ (op1(e12, e12) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f2282, plain, (~ spl20_10 | ~ spl20_12), inference(avatar_contradiction_clause, [], [f2281])).
fof(f2281, plain, ($false | (~ spl20_10 | ~ spl20_12)), inference(subsumption_resolution, [], [f2280, f260])).
fof(f2280, plain, ((e11 = e13) | (~ spl20_10 | ~ spl20_12)), inference(forward_demodulation, [], [f557, f549])).
fof(f549, plain, ((e11 = op1(e13, e11)) | ~ spl20_10), inference(avatar_component_clause, [], [f547])).
fof(f2272, plain, (~ spl20_5 | ~ spl20_6), inference(avatar_contradiction_clause, [], [f2271])).
fof(f2271, plain, ($false | (~ spl20_5 | ~ spl20_6)), inference(subsumption_resolution, [], [f2270, f256])).
fof(f256, plain, ~ (e10 = e11), inference(cnf_transformation, [], [f7])).
fof(f2270, plain, ((e10 = e11) | (~ spl20_5 | ~ spl20_6)), inference(backward_demodulation, [], [f532, f528])).
fof(f532, plain, ((e11 = op1(e13, e12)) | ~ spl20_6), inference(avatar_component_clause, [], [f530])).
fof(f530, plain, (spl20_6 <=> (e11 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_6])])).
fof(f2254, plain, (~ spl20_8 | ~ spl20_12), inference(avatar_split_clause, [], [f2252, f555, f538])).
fof(f2252, plain, (~ (e13 = op1(e13, e12)) | ~ spl20_12), inference(backward_demodulation, [], [f205, f557])).
fof(f205, plain, ~ (op1(e13, e11) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f2251, plain, (~ spl20_3 | ~ spl20_15), inference(avatar_split_clause, [], [f2248, f568, f517])).
fof(f517, plain, (spl20_3 <=> (e12 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_3])])).
fof(f2248, plain, (~ (e12 = op1(e13, e13)) | ~ spl20_15), inference(backward_demodulation, [], [f204, f570])).
fof(f204, plain, ~ (op1(e13, e10) = op1(e13, e13)), inference(cnf_transformation, [], [f5])).
fof(f2250, plain, (~ spl20_7 | ~ spl20_15), inference(avatar_split_clause, [], [f2247, f568, f534])).
fof(f534, plain, (spl20_7 <=> (e12 = op1(e13, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_7])])).
fof(f2247, plain, (~ (e12 = op1(e13, e12)) | ~ spl20_15), inference(backward_demodulation, [], [f203, f570])).
fof(f203, plain, ~ (op1(e13, e10) = op1(e13, e12)), inference(cnf_transformation, [], [f5])).
fof(f2249, plain, (~ spl20_11 | ~ spl20_15), inference(avatar_split_clause, [], [f2246, f568, f551])).
fof(f551, plain, (spl20_11 <=> (e12 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_11])])).
fof(f2246, plain, (~ (e12 = op1(e13, e11)) | ~ spl20_15), inference(backward_demodulation, [], [f202, f570])).
fof(f202, plain, ~ (op1(e13, e10) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f2244, plain, (~ spl20_19 | ~ spl20_20), inference(avatar_contradiction_clause, [], [f2243])).
fof(f2243, plain, ($false | (~ spl20_19 | ~ spl20_20)), inference(subsumption_resolution, [], [f2242, f261])).
fof(f261, plain, ~ (e12 = e13), inference(cnf_transformation, [], [f7])).
fof(f2242, plain, ((e12 = e13) | (~ spl20_19 | ~ spl20_20)), inference(backward_demodulation, [], [f591, f587])).
fof(f591, plain, ((e13 = op1(e12, e13)) | ~ spl20_20), inference(avatar_component_clause, [], [f589])).
fof(f2233, plain, (~ spl20_21 | ~ spl20_25), inference(avatar_split_clause, [], [f2230, f611, f594])).
fof(f2230, plain, (~ (e10 = op1(e12, e12)) | ~ spl20_25), inference(backward_demodulation, [], [f199, f613])).
fof(f199, plain, ~ (op1(e12, e11) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f2232, plain, (~ spl20_9 | ~ spl20_25), inference(avatar_split_clause, [], [f2229, f611, f543])).
fof(f543, plain, (spl20_9 <=> (e10 = op1(e13, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_9])])).
fof(f2229, plain, (~ (e10 = op1(e13, e11)) | ~ spl20_25), inference(backward_demodulation, [], [f171, f613])).
fof(f171, plain, ~ (op1(e12, e11) = op1(e13, e11)), inference(cnf_transformation, [], [f5])).
fof(f2227, plain, (~ spl20_18 | ~ spl20_30), inference(avatar_split_clause, [], [f2223, f632, f581])).
fof(f581, plain, (spl20_18 <=> (e11 = op1(e12, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_18])])).
fof(f2223, plain, (~ (e11 = op1(e12, e13)) | ~ spl20_30), inference(backward_demodulation, [], [f198, f634])).
fof(f198, plain, ~ (op1(e12, e10) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f2226, plain, (~ spl20_22 | ~ spl20_30), inference(avatar_split_clause, [], [f2222, f632, f598])).
fof(f598, plain, (spl20_22 <=> (e11 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_22])])).
fof(f2222, plain, (~ (e11 = op1(e12, e12)) | ~ spl20_30), inference(backward_demodulation, [], [f197, f634])).
fof(f197, plain, ~ (op1(e12, e10) = op1(e12, e12)), inference(cnf_transformation, [], [f5])).
fof(f2225, plain, (~ spl20_26 | ~ spl20_30), inference(avatar_split_clause, [], [f2221, f632, f615])).
fof(f615, plain, (spl20_26 <=> (e11 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_26])])).
fof(f2221, plain, (~ (e11 = op1(e12, e11)) | ~ spl20_30), inference(backward_demodulation, [], [f196, f634])).
fof(f196, plain, ~ (op1(e12, e10) = op1(e12, e11)), inference(cnf_transformation, [], [f5])).
fof(f2224, plain, (~ spl20_14 | ~ spl20_30), inference(avatar_split_clause, [], [f2220, f632, f564])).
fof(f564, plain, (spl20_14 <=> (e11 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_14])])).
fof(f2220, plain, (~ (e11 = op1(e13, e10)) | ~ spl20_30), inference(backward_demodulation, [], [f165, f634])).
fof(f165, plain, ~ (op1(e12, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f2185, plain, (~ spl20_37 | ~ spl20_45), inference(avatar_split_clause, [], [f2178, f696, f662])).
fof(f662, plain, (spl20_37 <=> (e10 = op1(e11, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_37])])).
fof(f2178, plain, (~ (e10 = op1(e11, e12)) | ~ spl20_45), inference(backward_demodulation, [], [f191, f698])).
fof(f191, plain, ~ (op1(e11, e10) = op1(e11, e12)), inference(cnf_transformation, [], [f5])).
fof(f2182, plain, (~ spl20_29 | ~ spl20_45), inference(avatar_split_clause, [], [f2175, f696, f628])).
fof(f628, plain, (spl20_29 <=> (e10 = op1(e12, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_29])])).
fof(f2175, plain, (~ (e10 = op1(e12, e10)) | ~ spl20_45), inference(backward_demodulation, [], [f163, f698])).
fof(f2167, plain, (~ spl20_19 | ~ spl20_51), inference(avatar_split_clause, [], [f2164, f721, f585])).
fof(f721, plain, (spl20_51 <=> (e12 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_51])])).
fof(f2164, plain, (~ (e12 = op1(e12, e13)) | ~ spl20_51), inference(backward_demodulation, [], [f179, f723])).
fof(f723, plain, ((e12 = op1(e10, e13)) | ~ spl20_51), inference(avatar_component_clause, [], [f721])).
fof(f179, plain, ~ (op1(e10, e13) = op1(e12, e13)), inference(cnf_transformation, [], [f5])).
fof(f2153, plain, (~ spl20_53 | ~ spl20_57), inference(avatar_split_clause, [], [f2148, f747, f730])).
fof(f2148, plain, (~ (e10 = op1(e10, e12)) | ~ spl20_57), inference(backward_demodulation, [], [f187, f749])).
fof(f749, plain, ((e10 = op1(e10, e11)) | ~ spl20_57), inference(avatar_component_clause, [], [f747])).
fof(f187, plain, ~ (op1(e10, e11) = op1(e10, e12)), inference(cnf_transformation, [], [f5])).
fof(f2144, plain, (spl20_15 | ~ spl20_64), inference(avatar_split_clause, [], [f2137, f776, f568])).
fof(f2137, plain, ((e12 = op1(e13, e10)) | ~ spl20_64), inference(backward_demodulation, [], [f423, f778])).
fof(f423, plain, (e12 = op1(op1(e10, e10), e10)), inference(cnf_transformation, [], [f12])).
fof(f12, plain, ((op1(e10, e10) = e13) & (e12 = op1(op1(e10, e10), e10)) & (e11 = op1(op1(op1(e10, e10), e10), e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax12)).
fof(f2143, plain, (~ spl20_52 | ~ spl20_64), inference(avatar_split_clause, [], [f2136, f776, f725])).
fof(f725, plain, (spl20_52 <=> (e13 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_52])])).
fof(f2136, plain, (~ (e13 = op1(e10, e13)) | ~ spl20_64), inference(backward_demodulation, [], [f186, f778])).
fof(f186, plain, ~ (op1(e10, e10) = op1(e10, e13)), inference(cnf_transformation, [], [f5])).
fof(f2141, plain, (~ spl20_60 | ~ spl20_64), inference(avatar_split_clause, [], [f2134, f776, f759])).
fof(f759, plain, (spl20_60 <=> (e13 = op1(e10, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_60])])).
fof(f2134, plain, (~ (e13 = op1(e10, e11)) | ~ spl20_64), inference(backward_demodulation, [], [f184, f778])).
fof(f184, plain, ~ (op1(e10, e10) = op1(e10, e11)), inference(cnf_transformation, [], [f5])).
fof(f2140, plain, (~ spl20_16 | ~ spl20_64), inference(avatar_split_clause, [], [f2133, f776, f572])).
fof(f572, plain, (spl20_16 <=> (e13 = op1(e13, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_16])])).
fof(f2133, plain, (~ (e13 = op1(e13, e10)) | ~ spl20_64), inference(backward_demodulation, [], [f162, f778])).
fof(f162, plain, ~ (op1(e10, e10) = op1(e13, e10)), inference(cnf_transformation, [], [f5])).
fof(f2138, plain, (~ spl20_48 | ~ spl20_64), inference(avatar_split_clause, [], [f2131, f776, f708])).
fof(f708, plain, (spl20_48 <=> (e13 = op1(e11, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_48])])).
fof(f2131, plain, (~ (e13 = op1(e11, e10)) | ~ spl20_64), inference(backward_demodulation, [], [f160, f778])).
fof(f160, plain, ~ (op1(e10, e10) = op1(e11, e10)), inference(cnf_transformation, [], [f5])).
fof(f2124, plain, (spl20_138 | ~ spl20_67), inference(avatar_split_clause, [], [f2123, f821, f1308])).
fof(f1308, plain, (spl20_138 <=> (e22 = h4(e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_138])])).
fof(f821, plain, (spl20_67 <=> (e22 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_67])])).
fof(f2123, plain, ((e22 = h4(e13)) | ~ spl20_67), inference(backward_demodulation, [], [f443, f823])).
fof(f823, plain, ((e22 = op2(e23, e23)) | ~ spl20_67), inference(avatar_component_clause, [], [f821])).
fof(f443, plain, (op2(e23, e23) = h4(e13)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((op2(e23, e23) = h4(e13)) & (op2(op2(e23, e23), e23) = h4(e12)) & (h4(e11) = op2(op2(op2(e23, e23), e23), e23)) & (e23 = h4(e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax17)).
fof(f2115, plain, (~ spl20_210 | ~ spl20_72), inference(avatar_split_clause, [], [f2113, f842, f1703])).
fof(f842, plain, (spl20_72 <=> (e23 = op2(e23, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_72])])).
fof(f2113, plain, (~ (e23 = h3(e13)) | ~ spl20_72), inference(backward_demodulation, [], [f1288, f844])).
fof(f844, plain, ((e23 = op2(e23, e22)) | ~ spl20_72), inference(avatar_component_clause, [], [f842])).
fof(f1288, plain, ~ (op2(e23, e22) = h3(e13)), inference(backward_demodulation, [], [f225, f439])).
fof(f225, plain, ~ (op2(e22, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f2104, plain, (~ spl20_229 | ~ spl20_76), inference(avatar_split_clause, [], [f2101, f859, f1813])).
fof(f2101, plain, (~ (e23 = h2(e13)) | ~ spl20_76), inference(backward_demodulation, [], [f1279, f861])).
fof(f1279, plain, ~ (op2(e23, e21) = h2(e13)), inference(backward_demodulation, [], [f218, f435])).
fof(f218, plain, ~ (op2(e21, e21) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f2098, plain, (~ spl20_75 | ~ spl20_79), inference(avatar_split_clause, [], [f2093, f872, f855])).
fof(f855, plain, (spl20_75 <=> (e22 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_75])])).
fof(f2093, plain, (~ (e22 = op2(e23, e21)) | ~ spl20_79), inference(backward_demodulation, [], [f250, f874])).
fof(f250, plain, ~ (op2(e23, e20) = op2(e23, e21)), inference(cnf_transformation, [], [f6])).
fof(f2095, plain, (~ spl20_138 | ~ spl20_79), inference(avatar_split_clause, [], [f2090, f872, f1308])).
fof(f2090, plain, (~ (e22 = h4(e13)) | ~ spl20_79), inference(backward_demodulation, [], [f1297, f874])).
fof(f1297, plain, ~ (op2(e23, e20) = h4(e13)), inference(backward_demodulation, [], [f252, f443])).
fof(f252, plain, ~ (op2(e23, e20) = op2(e23, e23)), inference(cnf_transformation, [], [f6])).
fof(f2089, plain, (~ spl20_83 | ~ spl20_84), inference(avatar_contradiction_clause, [], [f2088])).
fof(f2088, plain, ($false | (~ spl20_83 | ~ spl20_84)), inference(subsumption_resolution, [], [f2087, f267])).
fof(f267, plain, ~ (e22 = e23), inference(cnf_transformation, [], [f8])).
fof(f2087, plain, ((e22 = e23) | (~ spl20_83 | ~ spl20_84)), inference(backward_demodulation, [], [f895, f891])).
fof(f895, plain, ((e23 = op2(e22, e23)) | ~ spl20_84), inference(avatar_component_clause, [], [f893])).
fof(f893, plain, (spl20_84 <=> (e23 = op2(e22, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_84])])).
fof(f2079, plain, (spl20_158 | ~ spl20_85), inference(avatar_split_clause, [], [f2077, f898, f1409])).
fof(f898, plain, (spl20_85 <=> (e20 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_85])])).
fof(f2077, plain, ((e20 = h3(e13)) | ~ spl20_85), inference(backward_demodulation, [], [f439, f900])).
fof(f900, plain, ((e20 = op2(e22, e22)) | ~ spl20_85), inference(avatar_component_clause, [], [f898])).
fof(f2074, plain, (~ spl20_73 | ~ spl20_89), inference(avatar_split_clause, [], [f2069, f915, f847])).
fof(f847, plain, (spl20_73 <=> (e20 = op2(e23, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_73])])).
fof(f2069, plain, (~ (e20 = op2(e23, e21)) | ~ spl20_89), inference(backward_demodulation, [], [f219, f917])).
fof(f2064, plain, (spl20_178 | ~ spl20_94), inference(avatar_split_clause, [], [f2063, f936, f1513])).
fof(f2063, plain, ((e21 = h1(e11)) | ~ spl20_94), inference(backward_demodulation, [], [f1270, f938])).
fof(f2054, plain, (~ spl20_229 | ~ spl20_100), inference(avatar_split_clause, [], [f2051, f961, f1813])).
fof(f2051, plain, (~ (e23 = h2(e13)) | ~ spl20_100), inference(backward_demodulation, [], [f1282, f963])).
fof(f2041, plain, (spl20_170 | ~ spl20_105), inference(avatar_split_clause, [], [f2040, f983, f1470])).
fof(f983, plain, (spl20_105 <=> (e20 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_105])])).
fof(f2040, plain, ((e20 = h2(e13)) | ~ spl20_105), inference(backward_demodulation, [], [f435, f985])).
fof(f985, plain, ((e20 = op2(e21, e21)) | ~ spl20_105), inference(avatar_component_clause, [], [f983])).
fof(f2034, plain, (~ spl20_101 | ~ spl20_109), inference(avatar_split_clause, [], [f2027, f1000, f966])).
fof(f966, plain, (spl20_101 <=> (e20 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_101])])).
fof(f2027, plain, (~ (e20 = op2(e21, e22)) | ~ spl20_109), inference(backward_demodulation, [], [f239, f1002])).
fof(f239, plain, ~ (op2(e21, e20) = op2(e21, e22)), inference(cnf_transformation, [], [f6])).
fof(f2033, plain, (~ spl20_77 | ~ spl20_109), inference(avatar_split_clause, [], [f2026, f1000, f864])).
fof(f864, plain, (spl20_77 <=> (e20 = op2(e23, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_77])])).
fof(f2026, plain, (~ (e20 = op2(e23, e20)) | ~ spl20_109), inference(backward_demodulation, [], [f212, f1002])).
fof(f2019, plain, (~ spl20_83 | ~ spl20_115), inference(avatar_split_clause, [], [f2015, f1025, f889])).
fof(f1025, plain, (spl20_115 <=> (e22 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_115])])).
fof(f2015, plain, (~ (e22 = op2(e22, e23)) | ~ spl20_115), inference(backward_demodulation, [], [f227, f1027])).
fof(f1027, plain, ((e22 = op2(e20, e23)) | ~ spl20_115), inference(avatar_component_clause, [], [f1025])).
fof(f227, plain, ~ (op2(e20, e23) = op2(e22, e23)), inference(cnf_transformation, [], [f6])).
fof(f2010, plain, (~ spl20_69 | ~ spl20_117), inference(avatar_split_clause, [], [f2005, f1034, f830])).
fof(f2005, plain, (~ (e20 = op2(e23, e22)) | ~ spl20_117), inference(backward_demodulation, [], [f222, f1036])).
fof(f1036, plain, ((e20 = op2(e20, e22)) | ~ spl20_117), inference(avatar_component_clause, [], [f1034])).
fof(f222, plain, ~ (op2(e20, e22) = op2(e23, e22)), inference(cnf_transformation, [], [f6])).
fof(f1989, plain, (spl20_243 | ~ spl20_128), inference(avatar_split_clause, [], [f1988, f1080, f1911])).
fof(f1080, plain, (spl20_128 <=> (op2(e20, e20) = e23)), introduced(avatar_definition, [new_symbols(naming, [spl20_128])])).
fof(f1988, plain, ((e23 = h1(e13)) | ~ spl20_128), inference(backward_demodulation, [], [f431, f1082])).
fof(f1082, plain, ((op2(e20, e20) = e23) | ~ spl20_128), inference(avatar_component_clause, [], [f1080])).
fof(f1962, plain, (~ spl20_239 | ~ spl20_240 | ~ spl20_241 | ~ spl20_242 | spl20_179 | spl20_176 | spl20_173 | ~ spl20_243 | ~ spl20_244 | ~ spl20_245 | ~ spl20_246 | ~ spl20_247 | ~ spl20_248 | ~ spl20_249 | ~ spl20_250 | ~ spl20_251 | ~ spl20_252 | ~ spl20_253 | ~ spl20_254 | ~ spl20_255), inference(avatar_split_clause, [], [f1893, f1959, f1955, f1951, f1947, f1943, f1939, f1935, f1931, f1927, f1923, f1919, f1915, f1911, f1486, f1503, f1519, f1907, f1903, f1899, f1895])).
fof(f1519, plain, (spl20_179 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl20_179])])).
fof(f1503, plain, (spl20_176 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl20_176])])).
fof(f1486, plain, (spl20_173 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl20_173])])).
fof(f1893, plain, (~ (h1(e13) = h1(op1(e10, e10))) | ~ (h1(op1(e10, e11)) = op2(e20, h1(e11))) | ~ (op2(e20, e22) = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | ~ (h1(e11) = h1(op1(e12, e10))) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11)))), inference(forward_demodulation, [], [f1892, f431])).
fof(f1892, plain, (~ (op2(e20, e20) = h1(op1(e10, e10))) | ~ (h1(op1(e10, e11)) = op2(e20, h1(e11))) | ~ (op2(e20, e22) = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | ~ (h1(e11) = h1(op1(e12, e10))) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11)))), inference(forward_demodulation, [], [f1891, f428])).
fof(f1891, plain, (~ (h1(op1(e10, e11)) = op2(e20, h1(e11))) | ~ (op2(e20, e22) = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | ~ (h1(e11) = h1(op1(e12, e10))) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1890, f428])).
fof(f1890, plain, (~ (op2(e20, e22) = h1(op1(e10, e12))) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | ~ (h1(e11) = h1(op1(e12, e10))) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1889, f428])).
fof(f1889, plain, (~ (h1(op1(e10, e12)) = op2(h1(e10), e22)) | ~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | ~ (h1(e11) = h1(op1(e12, e10))) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1888, f1268])).
fof(f1888, plain, (~ (h1(op1(e10, e13)) = op2(e20, h1(e13))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | ~ (h1(e11) = h1(op1(e12, e10))) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1887, f428])).
fof(f1887, plain, (~ (h1(op1(e11, e10)) = op2(h1(e11), e20)) | ~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | ~ (h1(e11) = h1(op1(e12, e10))) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1886, f428])).
fof(f1886, plain, (~ (h1(op1(e11, e12)) = op2(h1(e11), e22)) | ~ (h1(e11) = h1(op1(e12, e10))) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1885, f1268])).
fof(f1885, plain, (~ (h1(e11) = h1(op1(e12, e10))) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1884, f1270])).
fof(f1884, plain, (~ (op2(e22, e20) = h1(op1(e12, e10))) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1883, f1268])).
fof(f1883, plain, (~ (h1(op1(e12, e10)) = op2(h1(e12), e20)) | ~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1882, f428])).
fof(f1882, plain, (~ (h1(op1(e12, e11)) = op2(e22, h1(e11))) | ~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1881, f1268])).
fof(f1881, plain, (~ (h3(e13) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1880, f439])).
fof(f1880, plain, (~ (op2(e22, e22) = h1(op1(e12, e12))) | ~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1879, f1268])).
fof(f1879, plain, (~ (h1(op1(e12, e13)) = op2(e22, h1(e13))) | ~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1878, f1268])).
fof(f1878, plain, (~ (e22 = h1(op1(e13, e10))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1877, f1266])).
fof(f1877, plain, (~ (h1(op1(e13, e10)) = op2(h1(e13), e20)) | ~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f1876, f428])).
fof(f1876, plain, (~ (h1(op1(e13, e12)) = op2(h1(e13), e22)) | ~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(forward_demodulation, [], [f495, f1268])).
fof(f495, plain, (~ (e23 = h1(e13)) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10)))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | sP19 | sP18 | sP17 | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | sP16 | sP15 | sP14 | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | sP13 | sP12 | sP11 | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | sP10 | sP9 | sP8 | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(definition_folding, [], [f20, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31])).
fof(f31, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP8), inference(usedef, [], [e31])).
fof(e31, plain, (sP8 <=> (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f32, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP9), inference(usedef, [], [e32])).
fof(e32, plain, (sP9 <=> (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f33, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP10), inference(usedef, [], [e33])).
fof(e33, plain, (sP10 <=> (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f34, plain, ((~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ sP11), inference(usedef, [], [e34])).
fof(e34, plain, (sP11 <=> (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f35, plain, ((~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | ~ sP12), inference(usedef, [], [e35])).
fof(e35, plain, (sP12 <=> (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f36, plain, ((~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | ~ sP13), inference(usedef, [], [e36])).
fof(e36, plain, (sP13 <=> (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f37, plain, ((~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ sP14), inference(usedef, [], [e37])).
fof(e37, plain, (sP14 <=> (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f38, plain, ((~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | ~ sP15), inference(usedef, [], [e38])).
fof(e38, plain, (sP15 <=> (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f39, plain, ((~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | ~ sP16), inference(usedef, [], [e39])).
fof(e39, plain, (sP16 <=> (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f40, plain, ((~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ sP17), inference(usedef, [], [e40])).
fof(e40, plain, (sP17 <=> (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f41, plain, ((~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | ~ sP18), inference(usedef, [], [e41])).
fof(e41, plain, (sP18 <=> (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f42, plain, ((~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | ~ sP19), inference(usedef, [], [e42])).
fof(e42, plain, (sP19 <=> (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f20, plain, (((~ (e23 = h4(e13)) & ~ (e23 = h4(e12)) & ~ (e23 = h4(e11)) & ~ (e23 = h4(e10))) | (~ (e22 = h4(e13)) & ~ (e22 = h4(e12)) & ~ (e22 = h4(e11)) & ~ (e22 = h4(e10))) | (~ (e21 = h4(e13)) & ~ (e21 = h4(e12)) & ~ (e21 = h4(e11)) & ~ (e21 = h4(e10))) | (~ (e20 = h4(e13)) & ~ (e20 = h4(e12)) & ~ (e20 = h4(e11)) & ~ (e20 = h4(e10))) | ~ (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) | ~ (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) | ~ (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) | ~ (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) | ~ (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) | ~ (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) | ~ (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) | ~ (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) | ~ (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) | ~ (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) | ~ (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) | ~ (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) | ~ (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) | ~ (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) | ~ (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) | ~ (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) & ((~ (e23 = h3(e13)) & ~ (e23 = h3(e12)) & ~ (e23 = h3(e11)) & ~ (e23 = h3(e10))) | (~ (e22 = h3(e13)) & ~ (e22 = h3(e12)) & ~ (e22 = h3(e11)) & ~ (e22 = h3(e10))) | (~ (e21 = h3(e13)) & ~ (e21 = h3(e12)) & ~ (e21 = h3(e11)) & ~ (e21 = h3(e10))) | (~ (e20 = h3(e13)) & ~ (e20 = h3(e12)) & ~ (e20 = h3(e11)) & ~ (e20 = h3(e10))) | ~ (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) | ~ (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) | ~ (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) | ~ (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) | ~ (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) | ~ (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) | ~ (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) | ~ (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) | ~ (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) | ~ (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) | ~ (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) | ~ (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) | ~ (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) | ~ (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) | ~ (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) | ~ (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) & ((~ (e23 = h2(e13)) & ~ (e23 = h2(e12)) & ~ (e23 = h2(e11)) & ~ (e23 = h2(e10))) | (~ (e22 = h2(e13)) & ~ (e22 = h2(e12)) & ~ (e22 = h2(e11)) & ~ (e22 = h2(e10))) | (~ (e21 = h2(e13)) & ~ (e21 = h2(e12)) & ~ (e21 = h2(e11)) & ~ (e21 = h2(e10))) | (~ (e20 = h2(e13)) & ~ (e20 = h2(e12)) & ~ (e20 = h2(e11)) & ~ (e20 = h2(e10))) | ~ (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) | ~ (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) | ~ (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) | ~ (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) | ~ (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) | ~ (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) | ~ (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) | ~ (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) | ~ (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) | ~ (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) | ~ (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) | ~ (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) | ~ (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) | ~ (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) | ~ (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) | ~ (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) & ((~ (e23 = h1(e13)) & ~ (e23 = h1(e12)) & ~ (e23 = h1(e11)) & ~ (e23 = h1(e10))) | (~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | (~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | (~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) | ~ (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) | ~ (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) | ~ (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) | ~ (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) | ~ (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) | ~ (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) | ~ (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) | ~ (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) | ~ (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) | ~ (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) | ~ (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) | ~ (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) | ~ (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) | ~ (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) | ~ (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), inference(negated_conjecture, [], [f18])).
fof(f18, plain, ~ ((((e23 = h4(e13)) | (e23 = h4(e12)) | (e23 = h4(e11)) | (e23 = h4(e10))) & ((e22 = h4(e13)) | (e22 = h4(e12)) | (e22 = h4(e11)) | (e22 = h4(e10))) & ((e21 = h4(e13)) | (e21 = h4(e12)) | (e21 = h4(e11)) | (e21 = h4(e10))) & ((e20 = h4(e13)) | (e20 = h4(e12)) | (e20 = h4(e11)) | (e20 = h4(e10))) & (h4(op1(e13, e13)) = op2(h4(e13), h4(e13))) & (h4(op1(e13, e12)) = op2(h4(e13), h4(e12))) & (h4(op1(e13, e11)) = op2(h4(e13), h4(e11))) & (h4(op1(e13, e10)) = op2(h4(e13), h4(e10))) & (h4(op1(e12, e13)) = op2(h4(e12), h4(e13))) & (h4(op1(e12, e12)) = op2(h4(e12), h4(e12))) & (h4(op1(e12, e11)) = op2(h4(e12), h4(e11))) & (h4(op1(e12, e10)) = op2(h4(e12), h4(e10))) & (h4(op1(e11, e13)) = op2(h4(e11), h4(e13))) & (h4(op1(e11, e12)) = op2(h4(e11), h4(e12))) & (h4(op1(e11, e11)) = op2(h4(e11), h4(e11))) & (h4(op1(e11, e10)) = op2(h4(e11), h4(e10))) & (h4(op1(e10, e13)) = op2(h4(e10), h4(e13))) & (h4(op1(e10, e12)) = op2(h4(e10), h4(e12))) & (h4(op1(e10, e11)) = op2(h4(e10), h4(e11))) & (h4(op1(e10, e10)) = op2(h4(e10), h4(e10)))) | (((e23 = h3(e13)) | (e23 = h3(e12)) | (e23 = h3(e11)) | (e23 = h3(e10))) & ((e22 = h3(e13)) | (e22 = h3(e12)) | (e22 = h3(e11)) | (e22 = h3(e10))) & ((e21 = h3(e13)) | (e21 = h3(e12)) | (e21 = h3(e11)) | (e21 = h3(e10))) & ((e20 = h3(e13)) | (e20 = h3(e12)) | (e20 = h3(e11)) | (e20 = h3(e10))) & (h3(op1(e13, e13)) = op2(h3(e13), h3(e13))) & (h3(op1(e13, e12)) = op2(h3(e13), h3(e12))) & (h3(op1(e13, e11)) = op2(h3(e13), h3(e11))) & (h3(op1(e13, e10)) = op2(h3(e13), h3(e10))) & (h3(op1(e12, e13)) = op2(h3(e12), h3(e13))) & (h3(op1(e12, e12)) = op2(h3(e12), h3(e12))) & (h3(op1(e12, e11)) = op2(h3(e12), h3(e11))) & (h3(op1(e12, e10)) = op2(h3(e12), h3(e10))) & (h3(op1(e11, e13)) = op2(h3(e11), h3(e13))) & (h3(op1(e11, e12)) = op2(h3(e11), h3(e12))) & (h3(op1(e11, e11)) = op2(h3(e11), h3(e11))) & (h3(op1(e11, e10)) = op2(h3(e11), h3(e10))) & (h3(op1(e10, e13)) = op2(h3(e10), h3(e13))) & (h3(op1(e10, e12)) = op2(h3(e10), h3(e12))) & (h3(op1(e10, e11)) = op2(h3(e10), h3(e11))) & (h3(op1(e10, e10)) = op2(h3(e10), h3(e10)))) | (((e23 = h2(e13)) | (e23 = h2(e12)) | (e23 = h2(e11)) | (e23 = h2(e10))) & ((e22 = h2(e13)) | (e22 = h2(e12)) | (e22 = h2(e11)) | (e22 = h2(e10))) & ((e21 = h2(e13)) | (e21 = h2(e12)) | (e21 = h2(e11)) | (e21 = h2(e10))) & ((e20 = h2(e13)) | (e20 = h2(e12)) | (e20 = h2(e11)) | (e20 = h2(e10))) & (h2(op1(e13, e13)) = op2(h2(e13), h2(e13))) & (h2(op1(e13, e12)) = op2(h2(e13), h2(e12))) & (h2(op1(e13, e11)) = op2(h2(e13), h2(e11))) & (h2(op1(e13, e10)) = op2(h2(e13), h2(e10))) & (h2(op1(e12, e13)) = op2(h2(e12), h2(e13))) & (h2(op1(e12, e12)) = op2(h2(e12), h2(e12))) & (h2(op1(e12, e11)) = op2(h2(e12), h2(e11))) & (h2(op1(e12, e10)) = op2(h2(e12), h2(e10))) & (h2(op1(e11, e13)) = op2(h2(e11), h2(e13))) & (h2(op1(e11, e12)) = op2(h2(e11), h2(e12))) & (h2(op1(e11, e11)) = op2(h2(e11), h2(e11))) & (h2(op1(e11, e10)) = op2(h2(e11), h2(e10))) & (h2(op1(e10, e13)) = op2(h2(e10), h2(e13))) & (h2(op1(e10, e12)) = op2(h2(e10), h2(e12))) & (h2(op1(e10, e11)) = op2(h2(e10), h2(e11))) & (h2(op1(e10, e10)) = op2(h2(e10), h2(e10)))) | (((e23 = h1(e13)) | (e23 = h1(e12)) | (e23 = h1(e11)) | (e23 = h1(e10))) & ((e22 = h1(e13)) | (e22 = h1(e12)) | (e22 = h1(e11)) | (e22 = h1(e10))) & ((e21 = h1(e13)) | (e21 = h1(e12)) | (e21 = h1(e11)) | (e21 = h1(e10))) & ((e20 = h1(e13)) | (e20 = h1(e12)) | (e20 = h1(e11)) | (e20 = h1(e10))) & (h1(op1(e13, e13)) = op2(h1(e13), h1(e13))) & (h1(op1(e13, e12)) = op2(h1(e13), h1(e12))) & (h1(op1(e13, e11)) = op2(h1(e13), h1(e11))) & (h1(op1(e13, e10)) = op2(h1(e13), h1(e10))) & (h1(op1(e12, e13)) = op2(h1(e12), h1(e13))) & (h1(op1(e12, e12)) = op2(h1(e12), h1(e12))) & (h1(op1(e12, e11)) = op2(h1(e12), h1(e11))) & (h1(op1(e12, e10)) = op2(h1(e12), h1(e10))) & (h1(op1(e11, e13)) = op2(h1(e11), h1(e13))) & (h1(op1(e11, e12)) = op2(h1(e11), h1(e12))) & (h1(op1(e11, e11)) = op2(h1(e11), h1(e11))) & (h1(op1(e11, e10)) = op2(h1(e11), h1(e10))) & (h1(op1(e10, e13)) = op2(h1(e10), h1(e13))) & (h1(op1(e10, e12)) = op2(h1(e10), h1(e12))) & (h1(op1(e10, e11)) = op2(h1(e10), h1(e11))) & (h1(op1(e10, e10)) = op2(h1(e10), h1(e10))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', co1)).
fof(f1534, plain, ~ spl20_179, inference(avatar_split_clause, [], [f1533, f1519])).
fof(f1533, plain, ~ sP8, inference(subsumption_resolution, [], [f488, f428])).
fof(f488, plain, (~ (e20 = h1(e10)) | ~ sP8), inference(cnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e20 = h1(e13)) & ~ (e20 = h1(e12)) & ~ (e20 = h1(e11)) & ~ (e20 = h1(e10))) | ~ sP8), inference(nnf_transformation, [], [f31])).
fof(f1516, plain, (~ spl20_176 | ~ spl20_178), inference(avatar_split_clause, [], [f485, f1513, f1503])).
fof(f485, plain, (~ (e21 = h1(e11)) | ~ sP9), inference(cnf_transformation, [], [f62])).
fof(f62, plain, ((~ (e21 = h1(e13)) & ~ (e21 = h1(e12)) & ~ (e21 = h1(e11)) & ~ (e21 = h1(e10))) | ~ sP9), inference(nnf_transformation, [], [f32])).
fof(f1495, plain, ~ spl20_173, inference(avatar_split_clause, [], [f1494, f1486])).
fof(f1494, plain, ~ sP10, inference(subsumption_resolution, [], [f482, f1268])).
fof(f482, plain, (~ (e22 = h1(e12)) | ~ sP10), inference(cnf_transformation, [], [f61])).
fof(f61, plain, ((~ (e22 = h1(e13)) & ~ (e22 = h1(e12)) & ~ (e22 = h1(e11)) & ~ (e22 = h1(e10))) | ~ sP10), inference(nnf_transformation, [], [f33])).
fof(f1259, plain, spl20_94, inference(avatar_split_clause, [], [f1258, f936])).
fof(f1258, plain, (e21 = op2(e22, e20)), inference(forward_demodulation, [], [f425, f426])).
fof(f425, plain, (e21 = op2(op2(op2(e20, e20), e20), e20)), inference(cnf_transformation, [], [f13])).
fof(f1257, plain, spl20_128, inference(avatar_split_clause, [], [f427, f1080])).
fof(f427, plain, (op2(e20, e20) = e23), inference(cnf_transformation, [], [f13])).
fof(f1256, plain, spl20_30, inference(avatar_split_clause, [], [f1255, f632])).
fof(f1255, plain, (e11 = op1(e12, e10)), inference(forward_demodulation, [], [f422, f423])).
fof(f422, plain, (e11 = op1(op1(op1(e10, e10), e10), e10)), inference(cnf_transformation, [], [f12])).
fof(f1254, plain, spl20_64, inference(avatar_split_clause, [], [f424, f776])).
fof(f424, plain, (op1(e10, e10) = e13), inference(cnf_transformation, [], [f12])).
fof(f1253, plain, ~ spl20_125, inference(avatar_split_clause, [], [f417, f1068])).
fof(f1068, plain, (spl20_125 <=> (e20 = op2(e20, e20))), introduced(avatar_definition, [new_symbols(naming, [spl20_125])])).
fof(f417, plain, ~ (e20 = op2(e20, e20)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ((sP7 | sP6 | sP5 | sP4) & ~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e20))), inference(definition_folding, [], [f11, e29, e28, e27, e26])).
fof(f26, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e23, e20)) | ~ (e20 = op2(e23, e22))) & ((e21 = op2(e23, e20)) | ~ (e20 = op2(e23, e21))) & ((e20 = op2(e23, e20)) | ~ (e20 = op2(e23, e20))) & ((e23 = op2(e22, e20)) | ~ (e20 = op2(e22, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e22, e20)) | ~ (e20 = op2(e22, e21))) & ((e20 = op2(e22, e20)) | ~ (e20 = op2(e22, e20))) & ((e23 = op2(e21, e20)) | ~ (e20 = op2(e21, e23))) & ((e22 = op2(e21, e20)) | ~ (e20 = op2(e21, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e21, e20)) | ~ (e20 = op2(e21, e20))) & ((op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & ((op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & ((op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)))) | ~ sP4), inference(usedef, [], [e26])).
fof(e26, plain, (sP4 <=> (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e23, e20)) | ~ (e20 = op2(e23, e22))) & ((e21 = op2(e23, e20)) | ~ (e20 = op2(e23, e21))) & ((e20 = op2(e23, e20)) | ~ (e20 = op2(e23, e20))) & ((e23 = op2(e22, e20)) | ~ (e20 = op2(e22, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e22, e20)) | ~ (e20 = op2(e22, e21))) & ((e20 = op2(e22, e20)) | ~ (e20 = op2(e22, e20))) & ((e23 = op2(e21, e20)) | ~ (e20 = op2(e21, e23))) & ((e22 = op2(e21, e20)) | ~ (e20 = op2(e21, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e21, e20)) | ~ (e20 = op2(e21, e20))) & ((op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & ((op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & ((op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f27, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e23, e21)) | ~ (e21 = op2(e23, e22))) & ((e21 = op2(e23, e21)) | ~ (e21 = op2(e23, e21))) & ((e20 = op2(e23, e21)) | ~ (e21 = op2(e23, e20))) & ((e23 = op2(e22, e21)) | ~ (e21 = op2(e22, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e22, e21)) | ~ (e21 = op2(e22, e21))) & ((e20 = op2(e22, e21)) | ~ (e21 = op2(e22, e20))) & ((e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & ((e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & ((e23 = op2(e20, e21)) | ~ (e21 = op2(e20, e23))) & ((e22 = op2(e20, e21)) | ~ (e21 = op2(e20, e22))) & ((e21 = op2(e20, e21)) | ~ (e21 = op2(e20, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21))) | ~ sP5), inference(usedef, [], [e27])).
fof(e27, plain, (sP5 <=> (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e23, e21)) | ~ (e21 = op2(e23, e22))) & ((e21 = op2(e23, e21)) | ~ (e21 = op2(e23, e21))) & ((e20 = op2(e23, e21)) | ~ (e21 = op2(e23, e20))) & ((e23 = op2(e22, e21)) | ~ (e21 = op2(e22, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e22, e21)) | ~ (e21 = op2(e22, e21))) & ((e20 = op2(e22, e21)) | ~ (e21 = op2(e22, e20))) & ((e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & ((e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & ((e23 = op2(e20, e21)) | ~ (e21 = op2(e20, e23))) & ((e22 = op2(e20, e21)) | ~ (e21 = op2(e20, e22))) & ((e21 = op2(e20, e21)) | ~ (e21 = op2(e20, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f28, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e23, e22)) | ~ (e22 = op2(e23, e22))) & ((e21 = op2(e23, e22)) | ~ (e22 = op2(e23, e21))) & ((e20 = op2(e23, e22)) | ~ (e22 = op2(e23, e20))) & ((e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & ((e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & ((e23 = op2(e21, e22)) | ~ (e22 = op2(e21, e23))) & ((e22 = op2(e21, e22)) | ~ (e22 = op2(e21, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e21, e22)) | ~ (e22 = op2(e21, e20))) & ((e23 = op2(e20, e22)) | ~ (e22 = op2(e20, e23))) & ((e22 = op2(e20, e22)) | ~ (e22 = op2(e20, e22))) & ((e21 = op2(e20, e22)) | ~ (e22 = op2(e20, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22))) | ~ sP6), inference(usedef, [], [e28])).
fof(e28, plain, (sP6 <=> (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e23, e22)) | ~ (e22 = op2(e23, e22))) & ((e21 = op2(e23, e22)) | ~ (e22 = op2(e23, e21))) & ((e20 = op2(e23, e22)) | ~ (e22 = op2(e23, e20))) & ((e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & ((e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & ((e23 = op2(e21, e22)) | ~ (e22 = op2(e21, e23))) & ((e22 = op2(e21, e22)) | ~ (e22 = op2(e21, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e21, e22)) | ~ (e22 = op2(e21, e20))) & ((e23 = op2(e20, e22)) | ~ (e22 = op2(e20, e23))) & ((e22 = op2(e20, e22)) | ~ (e22 = op2(e20, e22))) & ((e21 = op2(e20, e22)) | ~ (e22 = op2(e20, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f29, plain, ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & ((e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & ((e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | ~ (e23 = op2(e22, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e22, e23)) | ~ (e23 = op2(e22, e21))) & ((e20 = op2(e22, e23)) | ~ (e23 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | ~ (e23 = op2(e21, e23))) & ((e22 = op2(e21, e23)) | ~ (e23 = op2(e21, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e21, e23)) | ~ (e23 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | ~ (e23 = op2(e20, e23))) & ((e22 = op2(e20, e23)) | ~ (e23 = op2(e20, e22))) & ((e21 = op2(e20, e23)) | ~ (e23 = op2(e20, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23))) | ~ sP7), inference(usedef, [], [e29])).
fof(e29, plain, (sP7 <=> (((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & ((e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & ((e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | ~ (e23 = op2(e22, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e22, e23)) | ~ (e23 = op2(e22, e21))) & ((e20 = op2(e22, e23)) | ~ (e23 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | ~ (e23 = op2(e21, e23))) & ((e22 = op2(e21, e23)) | ~ (e23 = op2(e21, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e21, e23)) | ~ (e23 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | ~ (e23 = op2(e20, e23))) & ((e22 = op2(e20, e23)) | ~ (e23 = op2(e20, e22))) & ((e21 = op2(e20, e23)) | ~ (e23 = op2(e20, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f11, plain, (((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & ((e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & ((e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | ~ (e23 = op2(e22, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e22, e23)) | ~ (e23 = op2(e22, e21))) & ((e20 = op2(e22, e23)) | ~ (e23 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | ~ (e23 = op2(e21, e23))) & ((e22 = op2(e21, e23)) | ~ (e23 = op2(e21, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e21, e23)) | ~ (e23 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | ~ (e23 = op2(e20, e23))) & ((e22 = op2(e20, e23)) | ~ (e23 = op2(e20, e22))) & ((e21 = op2(e20, e23)) | ~ (e23 = op2(e20, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23))) | (((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e23, e22)) | ~ (e22 = op2(e23, e22))) & ((e21 = op2(e23, e22)) | ~ (e22 = op2(e23, e21))) & ((e20 = op2(e23, e22)) | ~ (e22 = op2(e23, e20))) & ((e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & ((e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & ((e23 = op2(e21, e22)) | ~ (e22 = op2(e21, e23))) & ((e22 = op2(e21, e22)) | ~ (e22 = op2(e21, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e21, e22)) | ~ (e22 = op2(e21, e20))) & ((e23 = op2(e20, e22)) | ~ (e22 = op2(e20, e23))) & ((e22 = op2(e20, e22)) | ~ (e22 = op2(e20, e22))) & ((e21 = op2(e20, e22)) | ~ (e22 = op2(e20, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22))) | (((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e23, e21)) | ~ (e21 = op2(e23, e22))) & ((e21 = op2(e23, e21)) | ~ (e21 = op2(e23, e21))) & ((e20 = op2(e23, e21)) | ~ (e21 = op2(e23, e20))) & ((e23 = op2(e22, e21)) | ~ (e21 = op2(e22, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e22, e21)) | ~ (e21 = op2(e22, e21))) & ((e20 = op2(e22, e21)) | ~ (e21 = op2(e22, e20))) & ((e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & ((e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & ((e23 = op2(e20, e21)) | ~ (e21 = op2(e20, e23))) & ((e22 = op2(e20, e21)) | ~ (e21 = op2(e20, e22))) & ((e21 = op2(e20, e21)) | ~ (e21 = op2(e20, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21))) | (((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e23, e20)) | ~ (e20 = op2(e23, e22))) & ((e21 = op2(e23, e20)) | ~ (e20 = op2(e23, e21))) & ((e20 = op2(e23, e20)) | ~ (e20 = op2(e23, e20))) & ((e23 = op2(e22, e20)) | ~ (e20 = op2(e22, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e22, e20)) | ~ (e20 = op2(e22, e21))) & ((e20 = op2(e22, e20)) | ~ (e20 = op2(e22, e20))) & ((e23 = op2(e21, e20)) | ~ (e20 = op2(e21, e23))) & ((e22 = op2(e21, e20)) | ~ (e20 = op2(e21, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e21, e20)) | ~ (e20 = op2(e21, e20))) & ((op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & ((op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & ((op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20))))) & ~ (e23 = op2(e23, e23)) & ~ (e22 = op2(e22, e22)) & ~ (e21 = op2(e21, e21)) & ~ (e20 = op2(e20, e20))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax11)).
fof(f1252, plain, ~ spl20_106, inference(avatar_split_clause, [], [f418, f987])).
fof(f987, plain, (spl20_106 <=> (e21 = op2(e21, e21))), introduced(avatar_definition, [new_symbols(naming, [spl20_106])])).
fof(f418, plain, ~ (e21 = op2(e21, e21)), inference(cnf_transformation, [], [f30])).
fof(f1251, plain, ~ spl20_87, inference(avatar_split_clause, [], [f419, f906])).
fof(f906, plain, (spl20_87 <=> (e22 = op2(e22, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_87])])).
fof(f419, plain, ~ (e22 = op2(e22, e22)), inference(cnf_transformation, [], [f30])).
fof(f1250, plain, ~ spl20_68, inference(avatar_split_clause, [], [f420, f825])).
fof(f825, plain, (spl20_68 <=> (e23 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_68])])).
fof(f420, plain, ~ (e23 = op2(e23, e23)), inference(cnf_transformation, [], [f30])).
fof(f1249, plain, (spl20_136 | spl20_135 | spl20_134 | spl20_133), inference(avatar_split_clause, [], [f421, f1186, f1202, f1218, f1234])).
fof(f1234, plain, (spl20_136 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl20_136])])).
fof(f1218, plain, (spl20_135 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl20_135])])).
fof(f1202, plain, (spl20_134 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl20_134])])).
fof(f1186, plain, (spl20_133 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl20_133])])).
fof(f421, plain, (sP7 | sP6 | sP5 | sP4), inference(cnf_transformation, [], [f30])).
fof(f1241, plain, (~ spl20_136 | ~ spl20_85 | spl20_95), inference(avatar_split_clause, [], [f411, f940, f898, f1234])).
fof(f411, plain, ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22)) | ~ sP4), inference(cnf_transformation, [], [f51])).
fof(f51, plain, ((((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23))) & ((e22 = op2(e23, e20)) | ~ (e20 = op2(e23, e22))) & ((e21 = op2(e23, e20)) | ~ (e20 = op2(e23, e21))) & ((e20 = op2(e23, e20)) | ~ (e20 = op2(e23, e20))) & ((e23 = op2(e22, e20)) | ~ (e20 = op2(e22, e23))) & ((e22 = op2(e22, e20)) | ~ (e20 = op2(e22, e22))) & ((e21 = op2(e22, e20)) | ~ (e20 = op2(e22, e21))) & ((e20 = op2(e22, e20)) | ~ (e20 = op2(e22, e20))) & ((e23 = op2(e21, e20)) | ~ (e20 = op2(e21, e23))) & ((e22 = op2(e21, e20)) | ~ (e20 = op2(e21, e22))) & ((e21 = op2(e21, e20)) | ~ (e20 = op2(e21, e21))) & ((e20 = op2(e21, e20)) | ~ (e20 = op2(e21, e20))) & ((op2(e20, e20) = e23) | ~ (e20 = op2(e20, e23))) & ((op2(e20, e20) = e22) | ~ (e20 = op2(e20, e22))) & ((op2(e20, e20) = e21) | ~ (e20 = op2(e20, e21))) & ((e20 = op2(e20, e20)) | ~ (e20 = op2(e20, e20)))) | ~ sP4), inference(nnf_transformation, [], [f26])).
fof(f1237, plain, (~ spl20_136 | ~ spl20_65 | spl20_80), inference(avatar_split_clause, [], [f416, f876, f813, f1234])).
fof(f813, plain, (spl20_65 <=> (e20 = op2(e23, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_65])])).
fof(f416, plain, ((e23 = op2(e23, e20)) | ~ (e20 = op2(e23, e23)) | ~ sP4), inference(cnf_transformation, [], [f51])).
fof(f1230, plain, (~ spl20_135 | ~ spl20_114 | spl20_124), inference(avatar_split_clause, [], [f388, f1063, f1021, f1218])).
fof(f1021, plain, (spl20_114 <=> (e21 = op2(e20, e23))), introduced(avatar_definition, [new_symbols(naming, [spl20_114])])).
fof(f388, plain, ((e23 = op2(e20, e21)) | ~ (e21 = op2(e20, e23)) | ~ sP5), inference(cnf_transformation, [], [f50])).
fof(f50, plain, ((((e23 = op2(e23, e21)) | ~ (e21 = op2(e23, e23))) & ((e22 = op2(e23, e21)) | ~ (e21 = op2(e23, e22))) & ((e21 = op2(e23, e21)) | ~ (e21 = op2(e23, e21))) & ((e20 = op2(e23, e21)) | ~ (e21 = op2(e23, e20))) & ((e23 = op2(e22, e21)) | ~ (e21 = op2(e22, e23))) & ((e22 = op2(e22, e21)) | ~ (e21 = op2(e22, e22))) & ((e21 = op2(e22, e21)) | ~ (e21 = op2(e22, e21))) & ((e20 = op2(e22, e21)) | ~ (e21 = op2(e22, e20))) & ((e23 = op2(e21, e21)) | ~ (e21 = op2(e21, e23))) & ((e22 = op2(e21, e21)) | ~ (e21 = op2(e21, e22))) & ((e21 = op2(e21, e21)) | ~ (e21 = op2(e21, e21))) & ((e20 = op2(e21, e21)) | ~ (e21 = op2(e21, e20))) & ((e23 = op2(e20, e21)) | ~ (e21 = op2(e20, e23))) & ((e22 = op2(e20, e21)) | ~ (e21 = op2(e20, e22))) & ((e21 = op2(e20, e21)) | ~ (e21 = op2(e20, e21))) & ((e20 = op2(e20, e21)) | ~ (op2(e20, e20) = e21))) | ~ sP5), inference(nnf_transformation, [], [f27])).
fof(f1226, plain, (~ spl20_135 | ~ spl20_94 | spl20_89), inference(avatar_split_clause, [], [f393, f915, f936, f1218])).
fof(f393, plain, ((e20 = op2(e22, e21)) | ~ (e21 = op2(e22, e20)) | ~ sP5), inference(cnf_transformation, [], [f50])).
fof(f1222, plain, (~ spl20_135 | ~ spl20_70 | spl20_75), inference(avatar_split_clause, [], [f399, f855, f834, f1218])).
fof(f399, plain, ((e22 = op2(e23, e21)) | ~ (e21 = op2(e23, e22)) | ~ sP5), inference(cnf_transformation, [], [f50])).
fof(f1212, plain, (~ spl20_134 | ~ spl20_107 | spl20_102), inference(avatar_split_clause, [], [f374, f970, f991, f1202])).
fof(f374, plain, ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21)) | ~ sP6), inference(cnf_transformation, [], [f49])).
fof(f49, plain, ((((e23 = op2(e23, e22)) | ~ (e22 = op2(e23, e23))) & ((e22 = op2(e23, e22)) | ~ (e22 = op2(e23, e22))) & ((e21 = op2(e23, e22)) | ~ (e22 = op2(e23, e21))) & ((e20 = op2(e23, e22)) | ~ (e22 = op2(e23, e20))) & ((e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23))) & ((e22 = op2(e22, e22)) | ~ (e22 = op2(e22, e22))) & ((e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21))) & ((e20 = op2(e22, e22)) | ~ (e22 = op2(e22, e20))) & ((e23 = op2(e21, e22)) | ~ (e22 = op2(e21, e23))) & ((e22 = op2(e21, e22)) | ~ (e22 = op2(e21, e22))) & ((e21 = op2(e21, e22)) | ~ (e22 = op2(e21, e21))) & ((e20 = op2(e21, e22)) | ~ (e22 = op2(e21, e20))) & ((e23 = op2(e20, e22)) | ~ (e22 = op2(e20, e23))) & ((e22 = op2(e20, e22)) | ~ (e22 = op2(e20, e22))) & ((e21 = op2(e20, e22)) | ~ (e22 = op2(e20, e21))) & ((e20 = op2(e20, e22)) | ~ (op2(e20, e20) = e22))) | ~ sP6), inference(nnf_transformation, [], [f28])).
fof(f1209, plain, (~ spl20_134 | ~ spl20_91 | spl20_86), inference(avatar_split_clause, [], [f378, f902, f923, f1202])).
fof(f378, plain, ((e21 = op2(e22, e22)) | ~ (e22 = op2(e22, e21)) | ~ sP6), inference(cnf_transformation, [], [f49])).
fof(f1208, plain, (~ spl20_134 | ~ spl20_83 | spl20_88), inference(avatar_split_clause, [], [f380, f910, f889, f1202])).
fof(f380, plain, ((e23 = op2(e22, e22)) | ~ (e22 = op2(e22, e23)) | ~ sP6), inference(cnf_transformation, [], [f49])).
fof(f1207, plain, (~ spl20_134 | ~ spl20_79 | spl20_69), inference(avatar_split_clause, [], [f381, f830, f872, f1202])).
fof(f381, plain, ((e20 = op2(e23, e22)) | ~ (e22 = op2(e23, e20)) | ~ sP6), inference(cnf_transformation, [], [f49])).
fof(f1200, plain, (~ spl20_133 | ~ spl20_128 | spl20_113), inference(avatar_split_clause, [], [f353, f1017, f1080, f1186])).
fof(f353, plain, ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23) | ~ sP7), inference(cnf_transformation, [], [f48])).
fof(f48, plain, ((((e23 = op2(e23, e23)) | ~ (e23 = op2(e23, e23))) & ((e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22))) & ((e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21))) & ((e20 = op2(e23, e23)) | ~ (e23 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | ~ (e23 = op2(e22, e23))) & ((e22 = op2(e22, e23)) | ~ (e23 = op2(e22, e22))) & ((e21 = op2(e22, e23)) | ~ (e23 = op2(e22, e21))) & ((e20 = op2(e22, e23)) | ~ (e23 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | ~ (e23 = op2(e21, e23))) & ((e22 = op2(e21, e23)) | ~ (e23 = op2(e21, e22))) & ((e21 = op2(e21, e23)) | ~ (e23 = op2(e21, e21))) & ((e20 = op2(e21, e23)) | ~ (e23 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | ~ (e23 = op2(e20, e23))) & ((e22 = op2(e20, e23)) | ~ (e23 = op2(e20, e22))) & ((e21 = op2(e20, e23)) | ~ (e23 = op2(e20, e21))) & ((e20 = op2(e20, e23)) | ~ (op2(e20, e20) = e23))) | ~ sP7), inference(nnf_transformation, [], [f29])).
fof(f1195, plain, (~ spl20_133 | ~ spl20_104 | spl20_99), inference(avatar_split_clause, [], [f359, f957, f978, f1186])).
fof(f978, plain, (spl20_104 <=> (e23 = op2(e21, e22))), introduced(avatar_definition, [new_symbols(naming, [spl20_104])])).
fof(f359, plain, ((e22 = op2(e21, e23)) | ~ (e23 = op2(e21, e22)) | ~ sP7), inference(cnf_transformation, [], [f48])).
fof(f1190, plain, (~ spl20_133 | ~ spl20_76 | spl20_66), inference(avatar_split_clause, [], [f366, f817, f859, f1186])).
fof(f366, plain, ((e21 = op2(e23, e23)) | ~ (e23 = op2(e23, e21)) | ~ sP7), inference(cnf_transformation, [], [f48])).
fof(f1189, plain, (~ spl20_133 | ~ spl20_72 | spl20_67), inference(avatar_split_clause, [], [f367, f821, f842, f1186])).
fof(f367, plain, ((e22 = op2(e23, e23)) | ~ (e23 = op2(e23, e22)) | ~ sP7), inference(cnf_transformation, [], [f48])).
fof(f1184, plain, ~ spl20_61, inference(avatar_split_clause, [], [f348, f764])).
fof(f764, plain, (spl20_61 <=> (e10 = op1(e10, e10))), introduced(avatar_definition, [new_symbols(naming, [spl20_61])])).
fof(f348, plain, ~ (e10 = op1(e10, e10)), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ((sP3 | sP2 | sP1 | sP0) & ~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e10))), inference(definition_folding, [], [f10, e24, e23, e22, e21])).
fof(f21, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e13, e10)) | ~ (e10 = op1(e13, e12))) & ((e11 = op1(e13, e10)) | ~ (e10 = op1(e13, e11))) & ((e10 = op1(e13, e10)) | ~ (e10 = op1(e13, e10))) & ((e13 = op1(e12, e10)) | ~ (e10 = op1(e12, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e12, e10)) | ~ (e10 = op1(e12, e11))) & ((e10 = op1(e12, e10)) | ~ (e10 = op1(e12, e10))) & ((e13 = op1(e11, e10)) | ~ (e10 = op1(e11, e13))) & ((e12 = op1(e11, e10)) | ~ (e10 = op1(e11, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e11, e10)) | ~ (e10 = op1(e11, e10))) & ((op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & ((op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & ((op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)))) | ~ sP0), inference(usedef, [], [e21])).
fof(e21, plain, (sP0 <=> (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e13, e10)) | ~ (e10 = op1(e13, e12))) & ((e11 = op1(e13, e10)) | ~ (e10 = op1(e13, e11))) & ((e10 = op1(e13, e10)) | ~ (e10 = op1(e13, e10))) & ((e13 = op1(e12, e10)) | ~ (e10 = op1(e12, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e12, e10)) | ~ (e10 = op1(e12, e11))) & ((e10 = op1(e12, e10)) | ~ (e10 = op1(e12, e10))) & ((e13 = op1(e11, e10)) | ~ (e10 = op1(e11, e13))) & ((e12 = op1(e11, e10)) | ~ (e10 = op1(e11, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e11, e10)) | ~ (e10 = op1(e11, e10))) & ((op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & ((op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & ((op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f22, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e13, e11)) | ~ (e11 = op1(e13, e12))) & ((e11 = op1(e13, e11)) | ~ (e11 = op1(e13, e11))) & ((e10 = op1(e13, e11)) | ~ (e11 = op1(e13, e10))) & ((e13 = op1(e12, e11)) | ~ (e11 = op1(e12, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e12, e11)) | ~ (e11 = op1(e12, e11))) & ((e10 = op1(e12, e11)) | ~ (e11 = op1(e12, e10))) & ((e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & ((e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & ((e13 = op1(e10, e11)) | ~ (e11 = op1(e10, e13))) & ((e12 = op1(e10, e11)) | ~ (e11 = op1(e10, e12))) & ((e11 = op1(e10, e11)) | ~ (e11 = op1(e10, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11))) | ~ sP1), inference(usedef, [], [e22])).
fof(e22, plain, (sP1 <=> (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e13, e11)) | ~ (e11 = op1(e13, e12))) & ((e11 = op1(e13, e11)) | ~ (e11 = op1(e13, e11))) & ((e10 = op1(e13, e11)) | ~ (e11 = op1(e13, e10))) & ((e13 = op1(e12, e11)) | ~ (e11 = op1(e12, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e12, e11)) | ~ (e11 = op1(e12, e11))) & ((e10 = op1(e12, e11)) | ~ (e11 = op1(e12, e10))) & ((e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & ((e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & ((e13 = op1(e10, e11)) | ~ (e11 = op1(e10, e13))) & ((e12 = op1(e10, e11)) | ~ (e11 = op1(e10, e12))) & ((e11 = op1(e10, e11)) | ~ (e11 = op1(e10, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f23, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e13, e12)) | ~ (e12 = op1(e13, e12))) & ((e11 = op1(e13, e12)) | ~ (e12 = op1(e13, e11))) & ((e10 = op1(e13, e12)) | ~ (e12 = op1(e13, e10))) & ((e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & ((e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & ((e13 = op1(e11, e12)) | ~ (e12 = op1(e11, e13))) & ((e12 = op1(e11, e12)) | ~ (e12 = op1(e11, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e11, e12)) | ~ (e12 = op1(e11, e10))) & ((e13 = op1(e10, e12)) | ~ (e12 = op1(e10, e13))) & ((e12 = op1(e10, e12)) | ~ (e12 = op1(e10, e12))) & ((e11 = op1(e10, e12)) | ~ (e12 = op1(e10, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12))) | ~ sP2), inference(usedef, [], [e23])).
fof(e23, plain, (sP2 <=> (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e13, e12)) | ~ (e12 = op1(e13, e12))) & ((e11 = op1(e13, e12)) | ~ (e12 = op1(e13, e11))) & ((e10 = op1(e13, e12)) | ~ (e12 = op1(e13, e10))) & ((e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & ((e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & ((e13 = op1(e11, e12)) | ~ (e12 = op1(e11, e13))) & ((e12 = op1(e11, e12)) | ~ (e12 = op1(e11, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e11, e12)) | ~ (e12 = op1(e11, e10))) & ((e13 = op1(e10, e12)) | ~ (e12 = op1(e10, e13))) & ((e12 = op1(e10, e12)) | ~ (e12 = op1(e10, e12))) & ((e11 = op1(e10, e12)) | ~ (e12 = op1(e10, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f24, plain, ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & ((e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & ((e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | ~ (e13 = op1(e12, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e12, e13)) | ~ (e13 = op1(e12, e11))) & ((e10 = op1(e12, e13)) | ~ (e13 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | ~ (e13 = op1(e11, e13))) & ((e12 = op1(e11, e13)) | ~ (e13 = op1(e11, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e11, e13)) | ~ (e13 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | ~ (e13 = op1(e10, e13))) & ((e12 = op1(e10, e13)) | ~ (e13 = op1(e10, e12))) & ((e11 = op1(e10, e13)) | ~ (e13 = op1(e10, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13))) | ~ sP3), inference(usedef, [], [e24])).
fof(e24, plain, (sP3 <=> (((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & ((e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & ((e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | ~ (e13 = op1(e12, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e12, e13)) | ~ (e13 = op1(e12, e11))) & ((e10 = op1(e12, e13)) | ~ (e13 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | ~ (e13 = op1(e11, e13))) & ((e12 = op1(e11, e13)) | ~ (e13 = op1(e11, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e11, e13)) | ~ (e13 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | ~ (e13 = op1(e10, e13))) & ((e12 = op1(e10, e13)) | ~ (e13 = op1(e10, e12))) & ((e11 = op1(e10, e13)) | ~ (e13 = op1(e10, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f10, plain, (((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & ((e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & ((e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | ~ (e13 = op1(e12, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e12, e13)) | ~ (e13 = op1(e12, e11))) & ((e10 = op1(e12, e13)) | ~ (e13 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | ~ (e13 = op1(e11, e13))) & ((e12 = op1(e11, e13)) | ~ (e13 = op1(e11, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e11, e13)) | ~ (e13 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | ~ (e13 = op1(e10, e13))) & ((e12 = op1(e10, e13)) | ~ (e13 = op1(e10, e12))) & ((e11 = op1(e10, e13)) | ~ (e13 = op1(e10, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13))) | (((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e13, e12)) | ~ (e12 = op1(e13, e12))) & ((e11 = op1(e13, e12)) | ~ (e12 = op1(e13, e11))) & ((e10 = op1(e13, e12)) | ~ (e12 = op1(e13, e10))) & ((e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & ((e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & ((e13 = op1(e11, e12)) | ~ (e12 = op1(e11, e13))) & ((e12 = op1(e11, e12)) | ~ (e12 = op1(e11, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e11, e12)) | ~ (e12 = op1(e11, e10))) & ((e13 = op1(e10, e12)) | ~ (e12 = op1(e10, e13))) & ((e12 = op1(e10, e12)) | ~ (e12 = op1(e10, e12))) & ((e11 = op1(e10, e12)) | ~ (e12 = op1(e10, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12))) | (((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e13, e11)) | ~ (e11 = op1(e13, e12))) & ((e11 = op1(e13, e11)) | ~ (e11 = op1(e13, e11))) & ((e10 = op1(e13, e11)) | ~ (e11 = op1(e13, e10))) & ((e13 = op1(e12, e11)) | ~ (e11 = op1(e12, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e12, e11)) | ~ (e11 = op1(e12, e11))) & ((e10 = op1(e12, e11)) | ~ (e11 = op1(e12, e10))) & ((e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & ((e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & ((e13 = op1(e10, e11)) | ~ (e11 = op1(e10, e13))) & ((e12 = op1(e10, e11)) | ~ (e11 = op1(e10, e12))) & ((e11 = op1(e10, e11)) | ~ (e11 = op1(e10, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11))) | (((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e13, e10)) | ~ (e10 = op1(e13, e12))) & ((e11 = op1(e13, e10)) | ~ (e10 = op1(e13, e11))) & ((e10 = op1(e13, e10)) | ~ (e10 = op1(e13, e10))) & ((e13 = op1(e12, e10)) | ~ (e10 = op1(e12, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e12, e10)) | ~ (e10 = op1(e12, e11))) & ((e10 = op1(e12, e10)) | ~ (e10 = op1(e12, e10))) & ((e13 = op1(e11, e10)) | ~ (e10 = op1(e11, e13))) & ((e12 = op1(e11, e10)) | ~ (e10 = op1(e11, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e11, e10)) | ~ (e10 = op1(e11, e10))) & ((op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & ((op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & ((op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10))))) & ~ (e13 = op1(e13, e13)) & ~ (e12 = op1(e12, e12)) & ~ (e11 = op1(e11, e11)) & ~ (e10 = op1(e10, e10))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax10)).
fof(f1183, plain, ~ spl20_42, inference(avatar_split_clause, [], [f349, f683])).
fof(f683, plain, (spl20_42 <=> (e11 = op1(e11, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_42])])).
fof(f349, plain, ~ (e11 = op1(e11, e11)), inference(cnf_transformation, [], [f25])).
fof(f1182, plain, ~ spl20_23, inference(avatar_split_clause, [], [f350, f602])).
fof(f602, plain, (spl20_23 <=> (e12 = op1(e12, e12))), introduced(avatar_definition, [new_symbols(naming, [spl20_23])])).
fof(f350, plain, ~ (e12 = op1(e12, e12)), inference(cnf_transformation, [], [f25])).
fof(f1181, plain, ~ spl20_4, inference(avatar_split_clause, [], [f351, f521])).
fof(f521, plain, (spl20_4 <=> (e13 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_4])])).
fof(f351, plain, ~ (e13 = op1(e13, e13)), inference(cnf_transformation, [], [f25])).
fof(f1180, plain, (spl20_132 | spl20_131 | spl20_130 | spl20_129), inference(avatar_split_clause, [], [f352, f1117, f1133, f1149, f1165])).
fof(f1165, plain, (spl20_132 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl20_132])])).
fof(f1149, plain, (spl20_131 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl20_131])])).
fof(f1133, plain, (spl20_130 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl20_130])])).
fof(f1117, plain, (spl20_129 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl20_129])])).
fof(f352, plain, (sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f25])).
fof(f1172, plain, (~ spl20_132 | ~ spl20_21 | spl20_31), inference(avatar_split_clause, [], [f342, f636, f594, f1165])).
fof(f342, plain, ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12)) | ~ sP0), inference(cnf_transformation, [], [f47])).
fof(f47, plain, ((((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13))) & ((e12 = op1(e13, e10)) | ~ (e10 = op1(e13, e12))) & ((e11 = op1(e13, e10)) | ~ (e10 = op1(e13, e11))) & ((e10 = op1(e13, e10)) | ~ (e10 = op1(e13, e10))) & ((e13 = op1(e12, e10)) | ~ (e10 = op1(e12, e13))) & ((e12 = op1(e12, e10)) | ~ (e10 = op1(e12, e12))) & ((e11 = op1(e12, e10)) | ~ (e10 = op1(e12, e11))) & ((e10 = op1(e12, e10)) | ~ (e10 = op1(e12, e10))) & ((e13 = op1(e11, e10)) | ~ (e10 = op1(e11, e13))) & ((e12 = op1(e11, e10)) | ~ (e10 = op1(e11, e12))) & ((e11 = op1(e11, e10)) | ~ (e10 = op1(e11, e11))) & ((e10 = op1(e11, e10)) | ~ (e10 = op1(e11, e10))) & ((op1(e10, e10) = e13) | ~ (e10 = op1(e10, e13))) & ((op1(e10, e10) = e12) | ~ (e10 = op1(e10, e12))) & ((op1(e10, e10) = e11) | ~ (e10 = op1(e10, e11))) & ((e10 = op1(e10, e10)) | ~ (e10 = op1(e10, e10)))) | ~ sP0), inference(nnf_transformation, [], [f21])).
fof(f1168, plain, (~ spl20_132 | ~ spl20_1 | spl20_16), inference(avatar_split_clause, [], [f347, f572, f509, f1165])).
fof(f509, plain, (spl20_1 <=> (e10 = op1(e13, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_1])])).
fof(f347, plain, ((e13 = op1(e13, e10)) | ~ (e10 = op1(e13, e13)) | ~ sP0), inference(cnf_transformation, [], [f47])).
fof(f1161, plain, (~ spl20_131 | ~ spl20_50 | spl20_60), inference(avatar_split_clause, [], [f319, f759, f717, f1149])).
fof(f717, plain, (spl20_50 <=> (e11 = op1(e10, e13))), introduced(avatar_definition, [new_symbols(naming, [spl20_50])])).
fof(f319, plain, ((e13 = op1(e10, e11)) | ~ (e11 = op1(e10, e13)) | ~ sP1), inference(cnf_transformation, [], [f46])).
fof(f46, plain, ((((e13 = op1(e13, e11)) | ~ (e11 = op1(e13, e13))) & ((e12 = op1(e13, e11)) | ~ (e11 = op1(e13, e12))) & ((e11 = op1(e13, e11)) | ~ (e11 = op1(e13, e11))) & ((e10 = op1(e13, e11)) | ~ (e11 = op1(e13, e10))) & ((e13 = op1(e12, e11)) | ~ (e11 = op1(e12, e13))) & ((e12 = op1(e12, e11)) | ~ (e11 = op1(e12, e12))) & ((e11 = op1(e12, e11)) | ~ (e11 = op1(e12, e11))) & ((e10 = op1(e12, e11)) | ~ (e11 = op1(e12, e10))) & ((e13 = op1(e11, e11)) | ~ (e11 = op1(e11, e13))) & ((e12 = op1(e11, e11)) | ~ (e11 = op1(e11, e12))) & ((e11 = op1(e11, e11)) | ~ (e11 = op1(e11, e11))) & ((e10 = op1(e11, e11)) | ~ (e11 = op1(e11, e10))) & ((e13 = op1(e10, e11)) | ~ (e11 = op1(e10, e13))) & ((e12 = op1(e10, e11)) | ~ (e11 = op1(e10, e12))) & ((e11 = op1(e10, e11)) | ~ (e11 = op1(e10, e11))) & ((e10 = op1(e10, e11)) | ~ (op1(e10, e10) = e11))) | ~ sP1), inference(nnf_transformation, [], [f22])).
fof(f1157, plain, (~ spl20_131 | ~ spl20_30 | spl20_25), inference(avatar_split_clause, [], [f324, f611, f632, f1149])).
fof(f324, plain, ((e10 = op1(e12, e11)) | ~ (e11 = op1(e12, e10)) | ~ sP1), inference(cnf_transformation, [], [f46])).
fof(f1153, plain, (~ spl20_131 | ~ spl20_6 | spl20_11), inference(avatar_split_clause, [], [f330, f551, f530, f1149])).
fof(f330, plain, ((e12 = op1(e13, e11)) | ~ (e11 = op1(e13, e12)) | ~ sP1), inference(cnf_transformation, [], [f46])).
fof(f1138, plain, (~ spl20_130 | ~ spl20_15 | spl20_5), inference(avatar_split_clause, [], [f312, f526, f568, f1133])).
fof(f312, plain, ((e10 = op1(e13, e12)) | ~ (e12 = op1(e13, e10)) | ~ sP2), inference(cnf_transformation, [], [f45])).
fof(f45, plain, ((((e13 = op1(e13, e12)) | ~ (e12 = op1(e13, e13))) & ((e12 = op1(e13, e12)) | ~ (e12 = op1(e13, e12))) & ((e11 = op1(e13, e12)) | ~ (e12 = op1(e13, e11))) & ((e10 = op1(e13, e12)) | ~ (e12 = op1(e13, e10))) & ((e13 = op1(e12, e12)) | ~ (e12 = op1(e12, e13))) & ((e12 = op1(e12, e12)) | ~ (e12 = op1(e12, e12))) & ((e11 = op1(e12, e12)) | ~ (e12 = op1(e12, e11))) & ((e10 = op1(e12, e12)) | ~ (e12 = op1(e12, e10))) & ((e13 = op1(e11, e12)) | ~ (e12 = op1(e11, e13))) & ((e12 = op1(e11, e12)) | ~ (e12 = op1(e11, e12))) & ((e11 = op1(e11, e12)) | ~ (e12 = op1(e11, e11))) & ((e10 = op1(e11, e12)) | ~ (e12 = op1(e11, e10))) & ((e13 = op1(e10, e12)) | ~ (e12 = op1(e10, e13))) & ((e12 = op1(e10, e12)) | ~ (e12 = op1(e10, e12))) & ((e11 = op1(e10, e12)) | ~ (e12 = op1(e10, e11))) & ((e10 = op1(e10, e12)) | ~ (op1(e10, e10) = e12))) | ~ sP2), inference(nnf_transformation, [], [f23])).
fof(f1131, plain, (~ spl20_129 | ~ spl20_64 | spl20_49), inference(avatar_split_clause, [], [f284, f713, f776, f1117])).
fof(f284, plain, ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13) | ~ sP3), inference(cnf_transformation, [], [f44])).
fof(f44, plain, ((((e13 = op1(e13, e13)) | ~ (e13 = op1(e13, e13))) & ((e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12))) & ((e11 = op1(e13, e13)) | ~ (e13 = op1(e13, e11))) & ((e10 = op1(e13, e13)) | ~ (e13 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | ~ (e13 = op1(e12, e13))) & ((e12 = op1(e12, e13)) | ~ (e13 = op1(e12, e12))) & ((e11 = op1(e12, e13)) | ~ (e13 = op1(e12, e11))) & ((e10 = op1(e12, e13)) | ~ (e13 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | ~ (e13 = op1(e11, e13))) & ((e12 = op1(e11, e13)) | ~ (e13 = op1(e11, e12))) & ((e11 = op1(e11, e13)) | ~ (e13 = op1(e11, e11))) & ((e10 = op1(e11, e13)) | ~ (e13 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | ~ (e13 = op1(e10, e13))) & ((e12 = op1(e10, e13)) | ~ (e13 = op1(e10, e12))) & ((e11 = op1(e10, e13)) | ~ (e13 = op1(e10, e11))) & ((e10 = op1(e10, e13)) | ~ (op1(e10, e10) = e13))) | ~ sP3), inference(nnf_transformation, [], [f24])).
fof(f1124, plain, (~ spl20_129 | ~ spl20_28 | spl20_18), inference(avatar_split_clause, [], [f293, f581, f623, f1117])).
fof(f623, plain, (spl20_28 <=> (e13 = op1(e12, e11))), introduced(avatar_definition, [new_symbols(naming, [spl20_28])])).
fof(f293, plain, ((e11 = op1(e12, e13)) | ~ (e13 = op1(e12, e11)) | ~ sP3), inference(cnf_transformation, [], [f44])).
fof(f1120, plain, (~ spl20_129 | ~ spl20_8 | spl20_3), inference(avatar_split_clause, [], [f298, f517, f538, f1117])).
fof(f298, plain, ((e12 = op1(e13, e13)) | ~ (e13 = op1(e13, e12)) | ~ sP3), inference(cnf_transformation, [], [f44])).
fof(f1115, plain, (spl20_125 | spl20_121 | spl20_117 | spl20_113), inference(avatar_split_clause, [], [f128, f1017, f1034, f1051, f1068])).
fof(f128, plain, ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20))), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))) & ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))) & ((e22 = op2(e23, e23)) | (e22 = op2(e22, e23)) | (e22 = op2(e21, e23)) | (e22 = op2(e20, e23))) & ((e22 = op2(e23, e23)) | (e22 = op2(e23, e22)) | (e22 = op2(e23, e21)) | (e22 = op2(e23, e20))) & ((e21 = op2(e23, e23)) | (e21 = op2(e22, e23)) | (e21 = op2(e21, e23)) | (e21 = op2(e20, e23))) & ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))) & ((e20 = op2(e23, e23)) | (e20 = op2(e22, e23)) | (e20 = op2(e21, e23)) | (e20 = op2(e20, e23))) & ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))) & ((e23 = op2(e23, e22)) | (e23 = op2(e22, e22)) | (e23 = op2(e21, e22)) | (e23 = op2(e20, e22))) & ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))) & ((e22 = op2(e23, e22)) | (e22 = op2(e22, e22)) | (e22 = op2(e21, e22)) | (e22 = op2(e20, e22))) & ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))) & ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))) & ((e21 = op2(e22, e23)) | (e21 = op2(e22, e22)) | (e21 = op2(e22, e21)) | (e21 = op2(e22, e20))) & ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))) & ((e20 = op2(e22, e23)) | (e20 = op2(e22, e22)) | (e20 = op2(e22, e21)) | (e20 = op2(e22, e20))) & ((e23 = op2(e23, e21)) | (e23 = op2(e22, e21)) | (e23 = op2(e21, e21)) | (e23 = op2(e20, e21))) & ((e23 = op2(e21, e23)) | (e23 = op2(e21, e22)) | (e23 = op2(e21, e21)) | (e23 = op2(e21, e20))) & ((e22 = op2(e23, e21)) | (e22 = op2(e22, e21)) | (e22 = op2(e21, e21)) | (e22 = op2(e20, e21))) & ((e22 = op2(e21, e23)) | (e22 = op2(e21, e22)) | (e22 = op2(e21, e21)) | (e22 = op2(e21, e20))) & ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))) & ((e21 = op2(e21, e23)) | (e21 = op2(e21, e22)) | (e21 = op2(e21, e21)) | (e21 = op2(e21, e20))) & ((e20 = op2(e23, e21)) | (e20 = op2(e22, e21)) | (e20 = op2(e21, e21)) | (e20 = op2(e20, e21))) & ((e20 = op2(e21, e23)) | (e20 = op2(e21, e22)) | (e20 = op2(e21, e21)) | (e20 = op2(e21, e20))) & ((e23 = op2(e23, e20)) | (e23 = op2(e22, e20)) | (e23 = op2(e21, e20)) | (op2(e20, e20) = e23)) & ((e23 = op2(e20, e23)) | (e23 = op2(e20, e22)) | (e23 = op2(e20, e21)) | (op2(e20, e20) = e23)) & ((e22 = op2(e23, e20)) | (e22 = op2(e22, e20)) | (e22 = op2(e21, e20)) | (op2(e20, e20) = e22)) & ((e22 = op2(e20, e23)) | (e22 = op2(e20, e22)) | (e22 = op2(e20, e21)) | (op2(e20, e20) = e22)) & ((e21 = op2(e23, e20)) | (e21 = op2(e22, e20)) | (e21 = op2(e21, e20)) | (op2(e20, e20) = e21)) & ((e21 = op2(e20, e23)) | (e21 = op2(e20, e22)) | (e21 = op2(e20, e21)) | (op2(e20, e20) = e21)) & ((e20 = op2(e23, e20)) | (e20 = op2(e22, e20)) | (e20 = op2(e21, e20)) | (e20 = op2(e20, e20))) & ((e20 = op2(e20, e23)) | (e20 = op2(e20, e22)) | (e20 = op2(e20, e21)) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax4)).
fof(f1104, plain, (spl20_122 | spl20_106 | spl20_90 | spl20_74), inference(avatar_split_clause, [], [f139, f851, f919, f987, f1055])).
fof(f139, plain, ((e21 = op2(e23, e21)) | (e21 = op2(e22, e21)) | (e21 = op2(e21, e21)) | (e21 = op2(e20, e21))), inference(cnf_transformation, [], [f4])).
fof(f1098, plain, (spl20_117 | spl20_101 | spl20_85 | spl20_69), inference(avatar_split_clause, [], [f145, f830, f898, f966, f1034])).
fof(f145, plain, ((e20 = op2(e23, e22)) | (e20 = op2(e22, e22)) | (e20 = op2(e21, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1096, plain, (spl20_118 | spl20_102 | spl20_86 | spl20_70), inference(avatar_split_clause, [], [f147, f834, f902, f970, f1038])).
fof(f147, plain, ((e21 = op2(e23, e22)) | (e21 = op2(e22, e22)) | (e21 = op2(e21, e22)) | (e21 = op2(e20, e22))), inference(cnf_transformation, [], [f4])).
fof(f1095, plain, (spl20_95 | spl20_91 | spl20_87 | spl20_83), inference(avatar_split_clause, [], [f148, f889, f906, f923, f940])).
fof(f148, plain, ((e22 = op2(e22, e23)) | (e22 = op2(e22, e22)) | (e22 = op2(e22, e21)) | (e22 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1093, plain, (spl20_96 | spl20_92 | spl20_88 | spl20_84), inference(avatar_split_clause, [], [f150, f893, f910, f927, f944])).
fof(f150, plain, ((e23 = op2(e22, e23)) | (e23 = op2(e22, e22)) | (e23 = op2(e22, e21)) | (e23 = op2(e22, e20))), inference(cnf_transformation, [], [f4])).
fof(f1091, plain, (spl20_77 | spl20_73 | spl20_69 | spl20_65), inference(avatar_split_clause, [], [f152, f813, f830, f847, f864])).
fof(f152, plain, ((e20 = op2(e23, e23)) | (e20 = op2(e23, e22)) | (e20 = op2(e23, e21)) | (e20 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1089, plain, (spl20_78 | spl20_74 | spl20_70 | spl20_66), inference(avatar_split_clause, [], [f154, f817, f834, f851, f868])).
fof(f154, plain, ((e21 = op2(e23, e23)) | (e21 = op2(e23, e22)) | (e21 = op2(e23, e21)) | (e21 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1085, plain, (spl20_80 | spl20_76 | spl20_72 | spl20_68), inference(avatar_split_clause, [], [f158, f825, f842, f859, f876])).
fof(f158, plain, ((e23 = op2(e23, e23)) | (e23 = op2(e23, e22)) | (e23 = op2(e23, e21)) | (e23 = op2(e23, e20))), inference(cnf_transformation, [], [f4])).
fof(f1084, plain, (spl20_116 | spl20_100 | spl20_84 | spl20_68), inference(avatar_split_clause, [], [f159, f825, f893, f961, f1029])).
fof(f159, plain, ((e23 = op2(e23, e23)) | (e23 = op2(e22, e23)) | (e23 = op2(e21, e23)) | (e23 = op2(e20, e23))), inference(cnf_transformation, [], [f4])).
fof(f1049, plain, (spl20_117 | spl20_118 | spl20_119 | spl20_120), inference(avatar_split_clause, [], [f114, f1046, f1042, f1038, f1034])).
fof(f114, plain, ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (((e23 = op2(e23, e23)) | (e22 = op2(e23, e23)) | (e21 = op2(e23, e23)) | (e20 = op2(e23, e23))) & ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))) & ((e23 = op2(e23, e21)) | (e22 = op2(e23, e21)) | (e21 = op2(e23, e21)) | (e20 = op2(e23, e21))) & ((e23 = op2(e23, e20)) | (e22 = op2(e23, e20)) | (e21 = op2(e23, e20)) | (e20 = op2(e23, e20))) & ((e23 = op2(e22, e23)) | (e22 = op2(e22, e23)) | (e21 = op2(e22, e23)) | (e20 = op2(e22, e23))) & ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))) & ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))) & ((e23 = op2(e22, e20)) | (e22 = op2(e22, e20)) | (e21 = op2(e22, e20)) | (e20 = op2(e22, e20))) & ((e23 = op2(e21, e23)) | (e22 = op2(e21, e23)) | (e21 = op2(e21, e23)) | (e20 = op2(e21, e23))) & ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))) & ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))) & ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))) & ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))) & ((e23 = op2(e20, e22)) | (e22 = op2(e20, e22)) | (e21 = op2(e20, e22)) | (e20 = op2(e20, e22))) & ((e23 = op2(e20, e21)) | (e22 = op2(e20, e21)) | (e21 = op2(e20, e21)) | (e20 = op2(e20, e21))) & ((op2(e20, e20) = e23) | (op2(e20, e20) = e22) | (op2(e20, e20) = e21) | (e20 = op2(e20, e20)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax3)).
fof(f1032, plain, (spl20_113 | spl20_114 | spl20_115 | spl20_116), inference(avatar_split_clause, [], [f115, f1029, f1025, f1021, f1017])).
fof(f115, plain, ((e23 = op2(e20, e23)) | (e22 = op2(e20, e23)) | (e21 = op2(e20, e23)) | (e20 = op2(e20, e23))), inference(cnf_transformation, [], [f3])).
fof(f1015, plain, (spl20_109 | spl20_110 | spl20_111 | spl20_112), inference(avatar_split_clause, [], [f116, f1012, f1008, f1004, f1000])).
fof(f116, plain, ((e23 = op2(e21, e20)) | (e22 = op2(e21, e20)) | (e21 = op2(e21, e20)) | (e20 = op2(e21, e20))), inference(cnf_transformation, [], [f3])).
fof(f998, plain, (spl20_105 | spl20_106 | spl20_107 | spl20_108), inference(avatar_split_clause, [], [f117, f995, f991, f987, f983])).
fof(f117, plain, ((e23 = op2(e21, e21)) | (e22 = op2(e21, e21)) | (e21 = op2(e21, e21)) | (e20 = op2(e21, e21))), inference(cnf_transformation, [], [f3])).
fof(f981, plain, (spl20_101 | spl20_102 | spl20_103 | spl20_104), inference(avatar_split_clause, [], [f118, f978, f974, f970, f966])).
fof(f118, plain, ((e23 = op2(e21, e22)) | (e22 = op2(e21, e22)) | (e21 = op2(e21, e22)) | (e20 = op2(e21, e22))), inference(cnf_transformation, [], [f3])).
fof(f930, plain, (spl20_89 | spl20_90 | spl20_91 | spl20_92), inference(avatar_split_clause, [], [f121, f927, f923, f919, f915])).
fof(f121, plain, ((e23 = op2(e22, e21)) | (e22 = op2(e22, e21)) | (e21 = op2(e22, e21)) | (e20 = op2(e22, e21))), inference(cnf_transformation, [], [f3])).
fof(f913, plain, (spl20_85 | spl20_86 | spl20_87 | spl20_88), inference(avatar_split_clause, [], [f122, f910, f906, f902, f898])).
fof(f122, plain, ((e23 = op2(e22, e22)) | (e22 = op2(e22, e22)) | (e21 = op2(e22, e22)) | (e20 = op2(e22, e22))), inference(cnf_transformation, [], [f3])).
fof(f845, plain, (spl20_69 | spl20_70 | spl20_71 | spl20_72), inference(avatar_split_clause, [], [f126, f842, f838, f834, f830])).
fof(f126, plain, ((e23 = op2(e23, e22)) | (e22 = op2(e23, e22)) | (e21 = op2(e23, e22)) | (e20 = op2(e23, e22))), inference(cnf_transformation, [], [f3])).
fof(f811, plain, (spl20_61 | spl20_57 | spl20_53 | spl20_49), inference(avatar_split_clause, [], [f80, f713, f730, f747, f764])).
fof(f80, plain, ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))) & ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))) & ((e12 = op1(e13, e13)) | (e12 = op1(e12, e13)) | (e12 = op1(e11, e13)) | (e12 = op1(e10, e13))) & ((e12 = op1(e13, e13)) | (e12 = op1(e13, e12)) | (e12 = op1(e13, e11)) | (e12 = op1(e13, e10))) & ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))) & ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))) & ((e10 = op1(e13, e13)) | (e10 = op1(e12, e13)) | (e10 = op1(e11, e13)) | (e10 = op1(e10, e13))) & ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))) & ((e13 = op1(e13, e12)) | (e13 = op1(e12, e12)) | (e13 = op1(e11, e12)) | (e13 = op1(e10, e12))) & ((e13 = op1(e12, e13)) | (e13 = op1(e12, e12)) | (e13 = op1(e12, e11)) | (e13 = op1(e12, e10))) & ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))) & ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))) & ((e11 = op1(e13, e12)) | (e11 = op1(e12, e12)) | (e11 = op1(e11, e12)) | (e11 = op1(e10, e12))) & ((e11 = op1(e12, e13)) | (e11 = op1(e12, e12)) | (e11 = op1(e12, e11)) | (e11 = op1(e12, e10))) & ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))) & ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))) & ((e13 = op1(e13, e11)) | (e13 = op1(e12, e11)) | (e13 = op1(e11, e11)) | (e13 = op1(e10, e11))) & ((e13 = op1(e11, e13)) | (e13 = op1(e11, e12)) | (e13 = op1(e11, e11)) | (e13 = op1(e11, e10))) & ((e12 = op1(e13, e11)) | (e12 = op1(e12, e11)) | (e12 = op1(e11, e11)) | (e12 = op1(e10, e11))) & ((e12 = op1(e11, e13)) | (e12 = op1(e11, e12)) | (e12 = op1(e11, e11)) | (e12 = op1(e11, e10))) & ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))) & ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))) & ((e10 = op1(e13, e11)) | (e10 = op1(e12, e11)) | (e10 = op1(e11, e11)) | (e10 = op1(e10, e11))) & ((e10 = op1(e11, e13)) | (e10 = op1(e11, e12)) | (e10 = op1(e11, e11)) | (e10 = op1(e11, e10))) & ((e13 = op1(e13, e10)) | (e13 = op1(e12, e10)) | (e13 = op1(e11, e10)) | (op1(e10, e10) = e13)) & ((e13 = op1(e10, e13)) | (e13 = op1(e10, e12)) | (e13 = op1(e10, e11)) | (op1(e10, e10) = e13)) & ((e12 = op1(e13, e10)) | (e12 = op1(e12, e10)) | (e12 = op1(e11, e10)) | (op1(e10, e10) = e12)) & ((e12 = op1(e10, e13)) | (e12 = op1(e10, e12)) | (e12 = op1(e10, e11)) | (op1(e10, e10) = e12)) & ((e11 = op1(e13, e10)) | (e11 = op1(e12, e10)) | (e11 = op1(e11, e10)) | (op1(e10, e10) = e11)) & ((e11 = op1(e10, e13)) | (e11 = op1(e10, e12)) | (e11 = op1(e10, e11)) | (op1(e10, e10) = e11)) & ((e10 = op1(e13, e10)) | (e10 = op1(e12, e10)) | (e10 = op1(e11, e10)) | (e10 = op1(e10, e10))) & ((e10 = op1(e10, e13)) | (e10 = op1(e10, e12)) | (e10 = op1(e10, e11)) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax2)).
fof(f801, plain, (spl20_46 | spl20_42 | spl20_38 | spl20_34), inference(avatar_split_clause, [], [f90, f649, f666, f683, f700])).
fof(f90, plain, ((e11 = op1(e11, e13)) | (e11 = op1(e11, e12)) | (e11 = op1(e11, e11)) | (e11 = op1(e11, e10))), inference(cnf_transformation, [], [f2])).
fof(f800, plain, (spl20_58 | spl20_42 | spl20_26 | spl20_10), inference(avatar_split_clause, [], [f91, f547, f615, f683, f751])).
fof(f91, plain, ((e11 = op1(e13, e11)) | (e11 = op1(e12, e11)) | (e11 = op1(e11, e11)) | (e11 = op1(e10, e11))), inference(cnf_transformation, [], [f2])).
fof(f795, plain, (spl20_29 | spl20_25 | spl20_21 | spl20_17), inference(avatar_split_clause, [], [f96, f577, f594, f611, f628])).
fof(f96, plain, ((e10 = op1(e12, e13)) | (e10 = op1(e12, e12)) | (e10 = op1(e12, e11)) | (e10 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f794, plain, (spl20_53 | spl20_37 | spl20_21 | spl20_5), inference(avatar_split_clause, [], [f97, f526, f594, f662, f730])).
fof(f97, plain, ((e10 = op1(e13, e12)) | (e10 = op1(e12, e12)) | (e10 = op1(e11, e12)) | (e10 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f791, plain, (spl20_31 | spl20_27 | spl20_23 | spl20_19), inference(avatar_split_clause, [], [f100, f585, f602, f619, f636])).
fof(f100, plain, ((e12 = op1(e12, e13)) | (e12 = op1(e12, e12)) | (e12 = op1(e12, e11)) | (e12 = op1(e12, e10))), inference(cnf_transformation, [], [f2])).
fof(f790, plain, (spl20_55 | spl20_39 | spl20_23 | spl20_7), inference(avatar_split_clause, [], [f101, f534, f602, f670, f738])).
fof(f101, plain, ((e12 = op1(e13, e12)) | (e12 = op1(e12, e12)) | (e12 = op1(e11, e12)) | (e12 = op1(e10, e12))), inference(cnf_transformation, [], [f2])).
fof(f787, plain, (spl20_13 | spl20_9 | spl20_5 | spl20_1), inference(avatar_split_clause, [], [f104, f509, f526, f543, f560])).
fof(f104, plain, ((e10 = op1(e13, e13)) | (e10 = op1(e13, e12)) | (e10 = op1(e13, e11)) | (e10 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f785, plain, (spl20_14 | spl20_10 | spl20_6 | spl20_2), inference(avatar_split_clause, [], [f106, f513, f530, f547, f564])).
fof(f106, plain, ((e11 = op1(e13, e13)) | (e11 = op1(e13, e12)) | (e11 = op1(e13, e11)) | (e11 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f784, plain, (spl20_50 | spl20_34 | spl20_18 | spl20_2), inference(avatar_split_clause, [], [f107, f513, f581, f649, f717])).
fof(f107, plain, ((e11 = op1(e13, e13)) | (e11 = op1(e12, e13)) | (e11 = op1(e11, e13)) | (e11 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f781, plain, (spl20_16 | spl20_12 | spl20_8 | spl20_4), inference(avatar_split_clause, [], [f110, f521, f538, f555, f572])).
fof(f110, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e13, e12)) | (e13 = op1(e13, e11)) | (e13 = op1(e13, e10))), inference(cnf_transformation, [], [f2])).
fof(f780, plain, (spl20_52 | spl20_36 | spl20_20 | spl20_4), inference(avatar_split_clause, [], [f111, f521, f589, f657, f725])).
fof(f111, plain, ((e13 = op1(e13, e13)) | (e13 = op1(e12, e13)) | (e13 = op1(e11, e13)) | (e13 = op1(e10, e13))), inference(cnf_transformation, [], [f2])).
fof(f762, plain, (spl20_57 | spl20_58 | spl20_59 | spl20_60), inference(avatar_split_clause, [], [f65, f759, f755, f751, f747])).
fof(f65, plain, ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e13 = op1(e13, e13)) | (e12 = op1(e13, e13)) | (e11 = op1(e13, e13)) | (e10 = op1(e13, e13))) & ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))) & ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))) & ((e13 = op1(e13, e10)) | (e12 = op1(e13, e10)) | (e11 = op1(e13, e10)) | (e10 = op1(e13, e10))) & ((e13 = op1(e12, e13)) | (e12 = op1(e12, e13)) | (e11 = op1(e12, e13)) | (e10 = op1(e12, e13))) & ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))) & ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))) & ((e13 = op1(e12, e10)) | (e12 = op1(e12, e10)) | (e11 = op1(e12, e10)) | (e10 = op1(e12, e10))) & ((e13 = op1(e11, e13)) | (e12 = op1(e11, e13)) | (e11 = op1(e11, e13)) | (e10 = op1(e11, e13))) & ((e13 = op1(e11, e12)) | (e12 = op1(e11, e12)) | (e11 = op1(e11, e12)) | (e10 = op1(e11, e12))) & ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))) & ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))) & ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))) & ((e13 = op1(e10, e12)) | (e12 = op1(e10, e12)) | (e11 = op1(e10, e12)) | (e10 = op1(e10, e12))) & ((e13 = op1(e10, e11)) | (e12 = op1(e10, e11)) | (e11 = op1(e10, e11)) | (e10 = op1(e10, e11))) & ((op1(e10, e10) = e13) | (op1(e10, e10) = e12) | (op1(e10, e10) = e11) | (e10 = op1(e10, e10)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG124+1.p', ax1)).
fof(f728, plain, (spl20_49 | spl20_50 | spl20_51 | spl20_52), inference(avatar_split_clause, [], [f67, f725, f721, f717, f713])).
fof(f67, plain, ((e13 = op1(e10, e13)) | (e12 = op1(e10, e13)) | (e11 = op1(e10, e13)) | (e10 = op1(e10, e13))), inference(cnf_transformation, [], [f1])).
fof(f711, plain, (spl20_45 | spl20_46 | spl20_47 | spl20_48), inference(avatar_split_clause, [], [f68, f708, f704, f700, f696])).
fof(f68, plain, ((e13 = op1(e11, e10)) | (e12 = op1(e11, e10)) | (e11 = op1(e11, e10)) | (e10 = op1(e11, e10))), inference(cnf_transformation, [], [f1])).
fof(f694, plain, (spl20_41 | spl20_42 | spl20_43 | spl20_44), inference(avatar_split_clause, [], [f69, f691, f687, f683, f679])).
fof(f69, plain, ((e13 = op1(e11, e11)) | (e12 = op1(e11, e11)) | (e11 = op1(e11, e11)) | (e10 = op1(e11, e11))), inference(cnf_transformation, [], [f1])).
fof(f626, plain, (spl20_25 | spl20_26 | spl20_27 | spl20_28), inference(avatar_split_clause, [], [f73, f623, f619, f615, f611])).
fof(f73, plain, ((e13 = op1(e12, e11)) | (e12 = op1(e12, e11)) | (e11 = op1(e12, e11)) | (e10 = op1(e12, e11))), inference(cnf_transformation, [], [f1])).
fof(f609, plain, (spl20_21 | spl20_22 | spl20_23 | spl20_24), inference(avatar_split_clause, [], [f74, f606, f602, f598, f594])).
fof(f74, plain, ((e13 = op1(e12, e12)) | (e12 = op1(e12, e12)) | (e11 = op1(e12, e12)) | (e10 = op1(e12, e12))), inference(cnf_transformation, [], [f1])).
fof(f558, plain, (spl20_9 | spl20_10 | spl20_11 | spl20_12), inference(avatar_split_clause, [], [f77, f555, f551, f547, f543])).
fof(f77, plain, ((e13 = op1(e13, e11)) | (e12 = op1(e13, e11)) | (e11 = op1(e13, e11)) | (e10 = op1(e13, e11))), inference(cnf_transformation, [], [f1])).
fof(f541, plain, (spl20_5 | spl20_6 | spl20_7 | spl20_8), inference(avatar_split_clause, [], [f78, f538, f534, f530, f526])).
fof(f78, plain, ((e13 = op1(e13, e12)) | (e12 = op1(e13, e12)) | (e11 = op1(e13, e12)) | (e10 = op1(e13, e12))), inference(cnf_transformation, [], [f1])).