fof(f3881, plain, $false, inference(avatar_sat_refutation, [], [f605, f622, f639, f656, f673, f690, f707, f724, f741, f758, f775, f792, f809, f826, f843, f860, f861, f862, f863, f864, f865, f866, f867, f868, f869, f870, f872, f873, f874, f875, f876, f877, f878, f879, f880, f882, f883, f885, f886, f887, f889, f890, f891, f892, f898, f899, f900, f907, f908, f914, f915, f916, f921, f922, f923, f924, f929, f930, f932, f937, f938, f939, f940, f945, f946, f947, f948, f954, f955, f956, f961, f962, f963, f964, f969, f970, f972, f977, f978, f979, f980, f986, f987, f988, f993, f994, f995, f996, f1001, f1002, f1003, f1004, f1009, f1010, f1017, f1018, f1020, f1026, f1027, f1033, f1034, f1035, f1036, f1041, f1042, f1043, f1044, f1050, f1051, f1052, f1059, f1060, f1067, f1068, f1073, f1074, f1075, f1076, f1081, f1082, f1083, f1084, f1090, f1091, f1092, f1097, f1098, f1105, f1106, f1107, f1108, f1113, f1114, f1115, f1116, f1122, f1123, f1124, f1129, f1130, f1131, f1132, f1137, f1138, f1145, f1146, f1153, f1154, f1155, f1156, f1162, f1163, f1169, f1170, f1171, f1172, f1177, f1178, f1179, f1180, f1185, f1186, f1194, f1195, f1196, f1201, f1202, f1203, f1204, f1210, f1211, f1212, f1219, f1220, f1227, f1228, f1234, f1235, f1236, f1241, f1242, f1243, f1244, f1249, f1250, f1251, f1252, f1258, f1259, f1265, f1266, f1273, f1274, f1276, f1281, f1282, f1283, f1284, f1289, f1290, f1291, f1292, f1298, f1299, f1305, f1306, f1307, f1308, f1313, f1314, f1316, f1321, f1322, f1323, f1324, f1330, f1331, f1337, f1338, f1339, f1340, f1345, f1346, f1347, f1348, f1353, f1354, f1362, f1363, f1371, f1372, f1377, f1379, f1380, f1387, f1388, f1395, f1396, f1401, f1402, f1403, f1404, f1409, f1410, f1411, f1412, f1417, f1418, f1419, f1420, f1423, f1424, f1425, f1426, f1427, f1428, f1437, f1446, f1451, f1456, f1457, f1458, f1467, f1486, f1488, f1497, f1506, f1511, f1516, f1517, f1518, f1541, f1547, f1558, f1560, f1563, f1566, f1573, f1583, f1595, f1604, f1618, f1619, f1649, f1650, f1656, f1657, f1658, f1661, f1664, f1699, f1702, f1708, f1716, f1724, f1726, f1741, f1743, f1758, f1767, f1777, f1833, f1837, f1853, f1858, f1884, f1921, f1953, f1955, f1957, f1983, f1992, f1995, f1999, f2007, f2033, f2070, f2074, f2090, f2110, f2111, f2113, f2134, f2161, f2189, f2201, f2212, f2220, f2236, f2240, f2248, f2274, f2298, f2314, f2333, f2339, f2344, f2345, f2356, f2369, f2382, f2388, f2391, f2395, f2408, f2412, f2439, f2443, f2459, f2472, f2479, f2493, f2539, f2553, f2556, f2560, f2566, f2568, f2635, f2640, f2653, f2656, f2661, f2667, f2685, f2690, f2694, f2733, f2736, f2743, f2747, f2757, f2761, f2767, f2772, f2774, f2787, f2798, f2799, f2811, f2813, f2817, f2820, f2822, f2824, f2835, f2838, f2852, f2853, f2877, f2916, f2950, f2954, f2955, f2967, f2984, f2985, f2993, f3003, f3005, f3027, f3051, f3057, f3060, f3067, f3068, f3073, f3076, f3077, f3080, f3088, f3089, f3106, f3116, f3151, f3153, f3160, f3190, f3209, f3220, f3221, f3222, f3234, f3238, f3246, f3254, f3277, f3292, f3309, f3320, f3321, f3322, f3330, f3340, f3347, f3358, f3366, f3367, f3376, f3385, f3394, f3411, f3416, f3422, f3427, f3435, f3436, f3444, f3445, f3449, f3464, f3479, f3486, f3491, f3506, f3535, f3536, f3539, f3543, f3548, f3549, f3551, f3553, f3572, f3575, f3577, f3585, f3587, f3590, f3593, f3597, f3600, f3612, f3639, f3640, f3647, f3654, f3672, f3677, f3685, f3692, f3695, f3727, f3747, f3748, f3761, f3766, f3769, f3778, f3794, f3795, f3803, f3816, f3817, f3826, f3827, f3837, f3838, f3843, f3848, f3876])).
fof(f3876, plain, (~ spl66_3 | ~ spl66_15), inference(avatar_split_clause, [], [f3874, f649, f598])).
fof(f598, plain, (spl66_3 <=> (e2 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_3])])).
fof(f649, plain, (spl66_15 <=> (e2 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_15])])).
fof(f3874, plain, (~ (e2 = op(e3, e3)) | ~ spl66_15), inference(backward_demodulation, [], [f279, f651])).
fof(f651, plain, ((e2 = op(e3, e0)) | ~ spl66_15), inference(avatar_component_clause, [], [f649])).
fof(f279, plain, ~ (op(e3, e0) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3, plain, (~ (op(e3, e2) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e3)) & ~ (op(e3, e1) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e3)) & ~ (op(e3, e0) = op(e3, e2)) & ~ (op(e3, e0) = op(e3, e1)) & ~ (op(e2, e2) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e3)) & ~ (op(e2, e1) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e3)) & ~ (op(e2, e0) = op(e2, e2)) & ~ (op(e2, e0) = op(e2, e1)) & ~ (op(e1, e2) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e3)) & ~ (op(e1, e1) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e3)) & ~ (op(e1, e0) = op(e1, e2)) & ~ (op(e1, e0) = op(e1, e1)) & ~ (op(e0, e2) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e3)) & ~ (op(e0, e1) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e3)) & ~ (op(e0, e0) = op(e0, e2)) & ~ (op(e0, e0) = op(e0, e1)) & ~ (op(e2, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e3, e3)) & ~ (op(e1, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e3, e3)) & ~ (op(e0, e3) = op(e2, e3)) & ~ (op(e0, e3) = op(e1, e3)) & ~ (op(e2, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e3, e2)) & ~ (op(e1, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e3, e2)) & ~ (op(e0, e2) = op(e2, e2)) & ~ (op(e0, e2) = op(e1, e2)) & ~ (op(e2, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e3, e1)) & ~ (op(e1, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e3, e1)) & ~ (op(e0, e1) = op(e2, e1)) & ~ (op(e0, e1) = op(e1, e1)) & ~ (op(e2, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e3, e0)) & ~ (op(e1, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e3, e0)) & ~ (op(e0, e0) = op(e2, e0)) & ~ (op(e0, e0) = op(e1, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax3)).
fof(f3848, plain, (~ spl66_16 | ~ spl66_32), inference(avatar_split_clause, [], [f3844, f721, f653])).
fof(f653, plain, (spl66_16 <=> (e3 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_16])])).
fof(f721, plain, (spl66_32 <=> (e3 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_32])])).
fof(f3844, plain, (~ (e3 = op(e3, e0)) | ~ spl66_32), inference(backward_demodulation, [], [f240, f723])).
fof(f723, plain, ((e3 = op(e2, e0)) | ~ spl66_32), inference(avatar_component_clause, [], [f721])).
fof(f240, plain, ~ (op(e2, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3843, plain, (~ spl66_20 | ~ spl66_36), inference(avatar_split_clause, [], [f3841, f738, f670])).
fof(f670, plain, (spl66_20 <=> (e3 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_20])])).
fof(f738, plain, (spl66_36 <=> (e3 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_36])])).
fof(f3841, plain, (~ (e3 = op(e2, e3)) | ~ spl66_36), inference(backward_demodulation, [], [f256, f740])).
fof(f740, plain, ((e3 = op(e1, e3)) | ~ spl66_36), inference(avatar_component_clause, [], [f738])).
fof(f256, plain, ~ (op(e1, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3838, plain, (~ spl66_33 | ~ spl66_37), inference(avatar_split_clause, [], [f3830, f743, f726])).
fof(f726, plain, (spl66_33 <=> (e0 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_33])])).
fof(f743, plain, (spl66_37 <=> (e0 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_37])])).
fof(f3830, plain, (~ (e0 = op(e1, e3)) | ~ spl66_37), inference(backward_demodulation, [], [f270, f745])).
fof(f745, plain, ((e0 = op(e1, e2)) | ~ spl66_37), inference(avatar_component_clause, [], [f743])).
fof(f270, plain, ~ (op(e1, e2) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3837, plain, (~ spl66_21 | ~ spl66_37), inference(avatar_split_clause, [], [f3828, f743, f675])).
fof(f675, plain, (spl66_21 <=> (e0 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_21])])).
fof(f3828, plain, (~ (e0 = op(e2, e2)) | ~ spl66_37), inference(backward_demodulation, [], [f250, f745])).
fof(f250, plain, ~ (op(e1, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3827, plain, (~ spl66_40 | ~ spl66_43 | spl66_139), inference(avatar_split_clause, [], [f3823, f1469, f768, f755])).
fof(f755, plain, (spl66_40 <=> (e3 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_40])])).
fof(f768, plain, (spl66_43 <=> (e2 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_43])])).
fof(f1469, plain, (spl66_139 <=> (e3 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_139])])).
fof(f3823, plain, (~ (e3 = op(e1, e2)) | (~ spl66_43 | spl66_139)), inference(backward_demodulation, [], [f1471, f770])).
fof(f770, plain, ((e2 = op(e1, e1)) | ~ spl66_43), inference(avatar_component_clause, [], [f768])).
fof(f1471, plain, (~ (e3 = op(e1, op(e1, e1))) | spl66_139), inference(avatar_component_clause, [], [f1469])).
fof(f3826, plain, (~ spl66_27 | ~ spl66_43), inference(avatar_split_clause, [], [f3818, f768, f700])).
fof(f700, plain, (spl66_27 <=> (e2 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_27])])).
fof(f3818, plain, (~ (e2 = op(e2, e1)) | ~ spl66_43), inference(backward_demodulation, [], [f244, f770])).
fof(f244, plain, ~ (op(e1, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3817, plain, (~ spl66_42 | ~ spl66_46), inference(avatar_split_clause, [], [f3811, f781, f764])).
fof(f764, plain, (spl66_42 <=> (e1 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_42])])).
fof(f781, plain, (spl66_46 <=> (e1 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_46])])).
fof(f3811, plain, (~ (e1 = op(e1, e1)) | ~ spl66_46), inference(backward_demodulation, [], [f265, f783])).
fof(f783, plain, ((e1 = op(e1, e0)) | ~ spl66_46), inference(avatar_component_clause, [], [f781])).
fof(f265, plain, ~ (op(e1, e0) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3816, plain, (~ spl66_30 | ~ spl66_46), inference(avatar_split_clause, [], [f3809, f781, f713])).
fof(f713, plain, (spl66_30 <=> (e1 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_30])])).
fof(f3809, plain, (~ (e1 = op(e2, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f238, f783])).
fof(f238, plain, ~ (op(e1, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3803, plain, (~ spl66_50 | ~ spl66_54), inference(avatar_split_clause, [], [f3799, f815, f798])).
fof(f798, plain, (spl66_50 <=> (e1 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_50])])).
fof(f815, plain, (spl66_54 <=> (e1 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_54])])).
fof(f3799, plain, (~ (e1 = op(e0, e3)) | ~ spl66_54), inference(backward_demodulation, [], [f264, f817])).
fof(f817, plain, ((e1 = op(e0, e2)) | ~ spl66_54), inference(avatar_component_clause, [], [f815])).
fof(f264, plain, ~ (op(e0, e2) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3795, plain, (~ spl66_50 | spl66_152 | ~ spl66_153), inference(avatar_split_clause, [], [f3792, f1538, f1533, f798])).
fof(f1533, plain, (spl66_152 <=> (e1 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_152])])).
fof(f1538, plain, (spl66_153 <=> (e0 = op(e3, op(e3, e3)))), introduced(avatar_definition, [new_symbols(naming, [spl66_153])])).
fof(f3792, plain, (~ (e1 = op(e0, e3)) | (spl66_152 | ~ spl66_153)), inference(backward_demodulation, [], [f1535, f1539])).
fof(f1539, plain, ((e0 = op(e3, op(e3, e3))) | ~ spl66_153), inference(avatar_component_clause, [], [f1538])).
fof(f1535, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | spl66_152), inference(avatar_component_clause, [], [f1533])).
fof(f3794, plain, (~ spl66_51 | spl66_150 | ~ spl66_153), inference(avatar_split_clause, [], [f3791, f1538, f1524, f802])).
fof(f802, plain, (spl66_51 <=> (e2 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_51])])).
fof(f1524, plain, (spl66_150 <=> (e2 = op(op(e3, op(e3, e3)), e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_150])])).
fof(f3791, plain, (~ (e2 = op(e0, e3)) | (spl66_150 | ~ spl66_153)), inference(backward_demodulation, [], [f1526, f1539])).
fof(f1526, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | spl66_150), inference(avatar_component_clause, [], [f1524])).
fof(f3778, plain, (~ spl66_2 | ~ spl66_6), inference(avatar_split_clause, [], [f3777, f611, f594])).
fof(f594, plain, (spl66_2 <=> (e1 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_2])])).
fof(f611, plain, (spl66_6 <=> (e1 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_6])])).
fof(f3777, plain, (~ (e1 = op(e3, e3)) | ~ spl66_6), inference(backward_demodulation, [], [f282, f613])).
fof(f613, plain, ((e1 = op(e3, e2)) | ~ spl66_6), inference(avatar_component_clause, [], [f611])).
fof(f282, plain, ~ (op(e3, e2) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3769, plain, (~ spl66_35 | ~ spl66_47), inference(avatar_split_clause, [], [f3767, f785, f734])).
fof(f734, plain, (spl66_35 <=> (e2 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_35])])).
fof(f785, plain, (spl66_47 <=> (e2 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_47])])).
fof(f3767, plain, (~ (e2 = op(e1, e3)) | ~ spl66_47), inference(backward_demodulation, [], [f267, f787])).
fof(f787, plain, ((e2 = op(e1, e0)) | ~ spl66_47), inference(avatar_component_clause, [], [f785])).
fof(f267, plain, ~ (op(e1, e0) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3766, plain, (~ spl66_2 | ~ spl66_50), inference(avatar_split_clause, [], [f3765, f798, f594])).
fof(f3765, plain, (~ (e1 = op(e3, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f255, f800])).
fof(f800, plain, ((e1 = op(e0, e3)) | ~ spl66_50), inference(avatar_component_clause, [], [f798])).
fof(f255, plain, ~ (op(e0, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3761, plain, (~ spl66_7 | ~ spl66_55), inference(avatar_split_clause, [], [f3758, f819, f615])).
fof(f615, plain, (spl66_7 <=> (e2 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_7])])).
fof(f819, plain, (spl66_55 <=> (e2 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_55])])).
fof(f3758, plain, (~ (e2 = op(e3, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f249, f821])).
fof(f821, plain, ((e2 = op(e0, e2)) | ~ spl66_55), inference(avatar_component_clause, [], [f819])).
fof(f249, plain, ~ (op(e0, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3748, plain, (~ spl66_49 | ~ spl66_61), inference(avatar_split_clause, [], [f3737, f845, f794])).
fof(f794, plain, (spl66_49 <=> (e0 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_49])])).
fof(f845, plain, (spl66_61 <=> (e0 = op(e0, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_61])])).
fof(f3737, plain, (~ (e0 = op(e0, e3)) | ~ spl66_61), inference(backward_demodulation, [], [f261, f847])).
fof(f847, plain, ((e0 = op(e0, e0)) | ~ spl66_61), inference(avatar_component_clause, [], [f845])).
fof(f261, plain, ~ (op(e0, e0) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3747, plain, (~ spl66_45 | ~ spl66_61), inference(avatar_split_clause, [], [f3735, f845, f777])).
fof(f777, plain, (spl66_45 <=> (e0 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_45])])).
fof(f3735, plain, (~ (e0 = op(e1, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f235, f847])).
fof(f235, plain, ~ (op(e0, e0) = op(e1, e0)), inference(cnf_transformation, [], [f3])).
fof(f3727, plain, (~ spl66_21 | ~ spl66_30 | ~ spl66_40 | spl66_144), inference(avatar_contradiction_clause, [], [f3726])).
fof(f3726, plain, ($false | (~ spl66_21 | ~ spl66_30 | ~ spl66_40 | spl66_144)), inference(subsumption_resolution, [], [f3725, f757])).
fof(f757, plain, ((e3 = op(e1, e2)) | ~ spl66_40), inference(avatar_component_clause, [], [f755])).
fof(f3725, plain, (~ (e3 = op(e1, e2)) | (~ spl66_21 | ~ spl66_30 | spl66_144)), inference(forward_demodulation, [], [f3724, f715])).
fof(f715, plain, ((e1 = op(e2, e0)) | ~ spl66_30), inference(avatar_component_clause, [], [f713])).
fof(f3724, plain, (~ (e3 = op(op(e2, e0), e2)) | (~ spl66_21 | spl66_144)), inference(forward_demodulation, [], [f1496, f677])).
fof(f677, plain, ((e0 = op(e2, e2)) | ~ spl66_21), inference(avatar_component_clause, [], [f675])).
fof(f1496, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | spl66_144), inference(avatar_component_clause, [], [f1494])).
fof(f1494, plain, (spl66_144 <=> (e3 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_144])])).
fof(f3695, plain, (~ spl66_3 | ~ spl66_7), inference(avatar_split_clause, [], [f3694, f615, f598])).
fof(f3694, plain, (~ (e2 = op(e3, e3)) | ~ spl66_7), inference(backward_demodulation, [], [f282, f617])).
fof(f617, plain, ((e2 = op(e3, e2)) | ~ spl66_7), inference(avatar_component_clause, [], [f615])).
fof(f3692, plain, (~ spl66_9 | ~ spl66_12), inference(avatar_contradiction_clause, [], [f3691])).
fof(f3691, plain, ($false | (~ spl66_9 | ~ spl66_12)), inference(subsumption_resolution, [], [f3690, f285])).
fof(f285, plain, ~ (e0 = e3), inference(cnf_transformation, [], [f4])).
fof(f4, plain, (~ (e2 = e3) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax4)).
fof(f3690, plain, ((e0 = e3) | (~ spl66_9 | ~ spl66_12)), inference(backward_demodulation, [], [f638, f626])).
fof(f626, plain, ((e0 = op(e3, e1)) | ~ spl66_9), inference(avatar_component_clause, [], [f624])).
fof(f624, plain, (spl66_9 <=> (e0 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_9])])).
fof(f638, plain, ((e3 = op(e3, e1)) | ~ spl66_12), inference(avatar_component_clause, [], [f636])).
fof(f636, plain, (spl66_12 <=> (e3 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_12])])).
fof(f3685, plain, (~ spl66_5 | ~ spl66_21), inference(avatar_split_clause, [], [f3680, f675, f607])).
fof(f607, plain, (spl66_5 <=> (e0 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_5])])).
fof(f3680, plain, (~ (e0 = op(e3, e2)) | ~ spl66_21), inference(backward_demodulation, [], [f252, f677])).
fof(f252, plain, ~ (op(e2, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3677, plain, (~ spl66_22 | ~ spl66_30), inference(avatar_split_clause, [], [f3674, f713, f679])).
fof(f679, plain, (spl66_22 <=> (e1 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_22])])).
fof(f3674, plain, (~ (e1 = op(e2, e2)) | ~ spl66_30), inference(backward_demodulation, [], [f272, f715])).
fof(f272, plain, ~ (op(e2, e0) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3672, plain, (~ spl66_3 | ~ spl66_35), inference(avatar_split_clause, [], [f3671, f734, f598])).
fof(f3671, plain, (~ (e2 = op(e3, e3)) | ~ spl66_35), inference(backward_demodulation, [], [f257, f736])).
fof(f736, plain, ((e2 = op(e1, e3)) | ~ spl66_35), inference(avatar_component_clause, [], [f734])).
fof(f257, plain, ~ (op(e1, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3654, plain, (~ spl66_33 | ~ spl66_49), inference(avatar_split_clause, [], [f3650, f794, f726])).
fof(f3650, plain, (~ (e0 = op(e1, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f253, f796])).
fof(f796, plain, ((e0 = op(e0, e3)) | ~ spl66_49), inference(avatar_component_clause, [], [f794])).
fof(f253, plain, ~ (op(e0, e3) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f3647, plain, (~ spl66_22 | ~ spl66_54), inference(avatar_split_clause, [], [f3642, f815, f679])).
fof(f3642, plain, (~ (e1 = op(e2, e2)) | ~ spl66_54), inference(backward_demodulation, [], [f248, f817])).
fof(f248, plain, ~ (op(e0, e2) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3640, plain, (~ spl66_55 | ~ spl66_63), inference(avatar_split_clause, [], [f3635, f853, f819])).
fof(f853, plain, (spl66_63 <=> (op(e0, e0) = e2)), introduced(avatar_definition, [new_symbols(naming, [spl66_63])])).
fof(f3635, plain, (~ (e2 = op(e0, e2)) | ~ spl66_63), inference(backward_demodulation, [], [f260, f855])).
fof(f855, plain, ((op(e0, e0) = e2) | ~ spl66_63), inference(avatar_component_clause, [], [f853])).
fof(f260, plain, ~ (op(e0, e0) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3639, plain, (~ spl66_31 | ~ spl66_63), inference(avatar_split_clause, [], [f3634, f853, f717])).
fof(f717, plain, (spl66_31 <=> (e2 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_31])])).
fof(f3634, plain, (~ (e2 = op(e2, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f236, f855])).
fof(f236, plain, ~ (op(e0, e0) = op(e2, e0)), inference(cnf_transformation, [], [f3])).
fof(f3612, plain, (~ spl66_52 | ~ spl66_60), inference(avatar_split_clause, [], [f3611, f840, f806])).
fof(f806, plain, (spl66_52 <=> (e3 = op(e0, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_52])])).
fof(f840, plain, (spl66_60 <=> (e3 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_60])])).
fof(f3611, plain, (~ (e3 = op(e0, e3)) | ~ spl66_60), inference(forward_demodulation, [], [f263, f842])).
fof(f842, plain, ((e3 = op(e0, e1)) | ~ spl66_60), inference(avatar_component_clause, [], [f840])).
fof(f263, plain, ~ (op(e0, e1) = op(e0, e3)), inference(cnf_transformation, [], [f3])).
fof(f3600, plain, (~ spl66_28 | ~ spl66_60), inference(avatar_split_clause, [], [f3599, f840, f704])).
fof(f704, plain, (spl66_28 <=> (e3 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_28])])).
fof(f3599, plain, (~ (e3 = op(e2, e1)) | ~ spl66_60), inference(forward_demodulation, [], [f242, f842])).
fof(f242, plain, ~ (op(e0, e1) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3597, plain, (~ spl66_56 | spl66_144 | ~ spl66_147), inference(avatar_split_clause, [], [f3596, f1508, f1494, f823])).
fof(f823, plain, (spl66_56 <=> (e3 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_56])])).
fof(f1508, plain, (spl66_147 <=> (e0 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_147])])).
fof(f3596, plain, (~ (e3 = op(e0, e2)) | (spl66_144 | ~ spl66_147)), inference(forward_demodulation, [], [f1496, f1509])).
fof(f1509, plain, ((e0 = op(e2, op(e2, e2))) | ~ spl66_147), inference(avatar_component_clause, [], [f1508])).
fof(f3593, plain, (~ spl66_24 | ~ spl66_40), inference(avatar_split_clause, [], [f3592, f755, f687])).
fof(f687, plain, (spl66_24 <=> (e3 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_24])])).
fof(f3592, plain, (~ (e3 = op(e2, e2)) | ~ spl66_40), inference(forward_demodulation, [], [f250, f757])).
fof(f3590, plain, (~ spl66_12 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f3589])).
fof(f3589, plain, ($false | (~ spl66_12 | ~ spl66_16)), inference(subsumption_resolution, [], [f3588, f655])).
fof(f655, plain, ((e3 = op(e3, e0)) | ~ spl66_16), inference(avatar_component_clause, [], [f653])).
fof(f3588, plain, (~ (e3 = op(e3, e0)) | ~ spl66_12), inference(forward_demodulation, [], [f277, f638])).
fof(f277, plain, ~ (op(e3, e0) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3587, plain, (~ spl66_44 | ~ spl66_12), inference(avatar_split_clause, [], [f3586, f636, f772])).
fof(f772, plain, (spl66_44 <=> (e3 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_44])])).
fof(f3586, plain, (~ (e3 = op(e1, e1)) | ~ spl66_12), inference(forward_demodulation, [], [f245, f638])).
fof(f245, plain, ~ (op(e1, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3585, plain, (~ spl66_12 | ~ spl66_60), inference(avatar_contradiction_clause, [], [f3584])).
fof(f3584, plain, ($false | (~ spl66_12 | ~ spl66_60)), inference(subsumption_resolution, [], [f3583, f842])).
fof(f3583, plain, (~ (e3 = op(e0, e1)) | ~ spl66_12), inference(forward_demodulation, [], [f243, f638])).
fof(f243, plain, ~ (op(e0, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f3577, plain, (~ spl66_8 | ~ spl66_16), inference(avatar_split_clause, [], [f3576, f653, f619])).
fof(f619, plain, (spl66_8 <=> (e3 = op(e3, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_8])])).
fof(f3576, plain, (~ (e3 = op(e3, e2)) | ~ spl66_16), inference(forward_demodulation, [], [f278, f655])).
fof(f278, plain, ~ (op(e3, e0) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3575, plain, (~ spl66_8 | ~ spl66_40), inference(avatar_split_clause, [], [f3574, f755, f619])).
fof(f3574, plain, (~ (e3 = op(e3, e2)) | ~ spl66_40), inference(forward_demodulation, [], [f251, f757])).
fof(f251, plain, ~ (op(e1, e2) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f3572, plain, (~ spl66_3 | ~ spl66_5 | spl66_153), inference(avatar_contradiction_clause, [], [f3571])).
fof(f3571, plain, ($false | (~ spl66_3 | ~ spl66_5 | spl66_153)), inference(subsumption_resolution, [], [f3562, f609])).
fof(f609, plain, ((e0 = op(e3, e2)) | ~ spl66_5), inference(avatar_component_clause, [], [f607])).
fof(f3562, plain, (~ (e0 = op(e3, e2)) | (~ spl66_3 | spl66_153)), inference(backward_demodulation, [], [f1540, f600])).
fof(f600, plain, ((e2 = op(e3, e3)) | ~ spl66_3), inference(avatar_component_clause, [], [f598])).
fof(f1540, plain, (~ (e0 = op(e3, op(e3, e3))) | spl66_153), inference(avatar_component_clause, [], [f1538])).
fof(f3553, plain, (~ spl66_4 | ~ spl66_16), inference(avatar_split_clause, [], [f3552, f653, f602])).
fof(f602, plain, (spl66_4 <=> (e3 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_4])])).
fof(f3552, plain, (~ (e3 = op(e3, e3)) | ~ spl66_16), inference(backward_demodulation, [], [f279, f655])).
fof(f3551, plain, (~ spl66_4 | ~ spl66_20), inference(avatar_split_clause, [], [f3550, f670, f602])).
fof(f3550, plain, (~ (e3 = op(e3, e3)) | ~ spl66_20), inference(backward_demodulation, [], [f258, f672])).
fof(f672, plain, ((e3 = op(e2, e3)) | ~ spl66_20), inference(avatar_component_clause, [], [f670])).
fof(f258, plain, ~ (op(e2, e3) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3549, plain, (~ spl66_19 | ~ spl66_31), inference(avatar_split_clause, [], [f3546, f717, f666])).
fof(f666, plain, (spl66_19 <=> (e2 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_19])])).
fof(f3546, plain, (~ (e2 = op(e2, e3)) | ~ spl66_31), inference(backward_demodulation, [], [f273, f719])).
fof(f719, plain, ((e2 = op(e2, e0)) | ~ spl66_31), inference(avatar_component_clause, [], [f717])).
fof(f273, plain, ~ (op(e2, e0) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3548, plain, (~ spl66_15 | ~ spl66_31), inference(avatar_split_clause, [], [f3545, f717, f649])).
fof(f3545, plain, (~ (e2 = op(e3, e0)) | ~ spl66_31), inference(backward_demodulation, [], [f240, f719])).
fof(f3543, plain, (~ spl66_39 | ~ spl66_43), inference(avatar_split_clause, [], [f3540, f768, f751])).
fof(f751, plain, (spl66_39 <=> (e2 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_39])])).
fof(f3540, plain, (~ (e2 = op(e1, e2)) | ~ spl66_43), inference(backward_demodulation, [], [f268, f770])).
fof(f268, plain, ~ (op(e1, e1) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3539, plain, (~ spl66_39 | ~ spl66_55), inference(avatar_split_clause, [], [f3537, f819, f751])).
fof(f3537, plain, (~ (e2 = op(e1, e2)) | ~ spl66_55), inference(backward_demodulation, [], [f247, f821])).
fof(f247, plain, ~ (op(e0, e2) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3536, plain, (~ spl66_56 | ~ spl66_60), inference(avatar_split_clause, [], [f3534, f840, f823])).
fof(f3534, plain, (~ (e3 = op(e0, e2)) | ~ spl66_60), inference(backward_demodulation, [], [f262, f842])).
fof(f262, plain, ~ (op(e0, e1) = op(e0, e2)), inference(cnf_transformation, [], [f3])).
fof(f3535, plain, (~ spl66_44 | ~ spl66_60), inference(avatar_split_clause, [], [f3533, f840, f772])).
fof(f3533, plain, (~ (e3 = op(e1, e1)) | ~ spl66_60), inference(backward_demodulation, [], [f241, f842])).
fof(f241, plain, ~ (op(e0, e1) = op(e1, e1)), inference(cnf_transformation, [], [f3])).
fof(f3506, plain, (~ spl66_56 | ~ spl66_22 | ~ spl66_25 | spl66_144), inference(avatar_split_clause, [], [f3481, f1494, f692, f679, f823])).
fof(f692, plain, (spl66_25 <=> (e0 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_25])])).
fof(f3481, plain, (~ (e3 = op(e0, e2)) | (~ spl66_22 | ~ spl66_25 | spl66_144)), inference(forward_demodulation, [], [f3475, f694])).
fof(f694, plain, ((e0 = op(e2, e1)) | ~ spl66_25), inference(avatar_component_clause, [], [f692])).
fof(f3475, plain, (~ (e3 = op(op(e2, e1), e2)) | (~ spl66_22 | spl66_144)), inference(backward_demodulation, [], [f1496, f681])).
fof(f681, plain, ((e1 = op(e2, e2)) | ~ spl66_22), inference(avatar_component_clause, [], [f679])).
fof(f3491, plain, (~ spl66_11 | ~ spl66_15), inference(avatar_split_clause, [], [f3489, f649, f632])).
fof(f632, plain, (spl66_11 <=> (e2 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_11])])).
fof(f3489, plain, (~ (e2 = op(e3, e1)) | ~ spl66_15), inference(backward_demodulation, [], [f277, f651])).
fof(f3486, plain, (~ spl66_22 | ~ spl66_25 | spl66_147), inference(avatar_contradiction_clause, [], [f3485])).
fof(f3485, plain, ($false | (~ spl66_22 | ~ spl66_25 | spl66_147)), inference(subsumption_resolution, [], [f3477, f694])).
fof(f3477, plain, (~ (e0 = op(e2, e1)) | (~ spl66_22 | spl66_147)), inference(backward_demodulation, [], [f1510, f681])).
fof(f1510, plain, (~ (e0 = op(e2, op(e2, e2))) | spl66_147), inference(avatar_component_clause, [], [f1508])).
fof(f3479, plain, (~ spl66_18 | ~ spl66_22), inference(avatar_split_clause, [], [f3471, f679, f662])).
fof(f662, plain, (spl66_18 <=> (e1 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_18])])).
fof(f3471, plain, (~ (e1 = op(e2, e3)) | ~ spl66_22), inference(backward_demodulation, [], [f276, f681])).
fof(f276, plain, ~ (op(e2, e2) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3464, plain, (~ spl66_28 | ~ spl66_32), inference(avatar_split_clause, [], [f3461, f721, f704])).
fof(f3461, plain, (~ (e3 = op(e2, e1)) | ~ spl66_32), inference(backward_demodulation, [], [f271, f723])).
fof(f271, plain, ~ (op(e2, e0) = op(e2, e1)), inference(cnf_transformation, [], [f3])).
fof(f3449, plain, (~ spl66_23 | ~ spl66_39), inference(avatar_split_clause, [], [f3446, f751, f683])).
fof(f683, plain, (spl66_23 <=> (e2 = op(e2, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_23])])).
fof(f3446, plain, (~ (e2 = op(e2, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f250, f753])).
fof(f753, plain, ((e2 = op(e1, e2)) | ~ spl66_39), inference(avatar_component_clause, [], [f751])).
fof(f3445, plain, (~ spl66_35 | ~ spl66_44 | spl66_137), inference(avatar_split_clause, [], [f3441, f1460, f772, f734])).
fof(f1460, plain, (spl66_137 <=> (e2 = op(e1, op(e1, e1)))), introduced(avatar_definition, [new_symbols(naming, [spl66_137])])).
fof(f3441, plain, (~ (e2 = op(e1, e3)) | (~ spl66_44 | spl66_137)), inference(backward_demodulation, [], [f1462, f774])).
fof(f774, plain, ((e3 = op(e1, e1)) | ~ spl66_44), inference(avatar_component_clause, [], [f772])).
fof(f1462, plain, (~ (e2 = op(e1, op(e1, e1))) | spl66_137), inference(avatar_component_clause, [], [f1460])).
fof(f3444, plain, (~ spl66_28 | ~ spl66_44), inference(avatar_split_clause, [], [f3437, f772, f704])).
fof(f3437, plain, (~ (e3 = op(e2, e1)) | ~ spl66_44), inference(backward_demodulation, [], [f244, f774])).
fof(f3436, plain, (~ spl66_38 | ~ spl66_46), inference(avatar_split_clause, [], [f3431, f781, f747])).
fof(f747, plain, (spl66_38 <=> (e1 = op(e1, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_38])])).
fof(f3431, plain, (~ (e1 = op(e1, e2)) | ~ spl66_46), inference(backward_demodulation, [], [f266, f783])).
fof(f266, plain, ~ (op(e1, e0) = op(e1, e2)), inference(cnf_transformation, [], [f3])).
fof(f3435, plain, (~ spl66_14 | ~ spl66_46), inference(avatar_split_clause, [], [f3429, f781, f645])).
fof(f645, plain, (spl66_14 <=> (e1 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_14])])).
fof(f3429, plain, (~ (e1 = op(e3, e0)) | ~ spl66_46), inference(backward_demodulation, [], [f239, f783])).
fof(f239, plain, ~ (op(e1, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3427, plain, (~ spl66_18 | ~ spl66_50), inference(avatar_split_clause, [], [f3424, f798, f662])).
fof(f3424, plain, (~ (e1 = op(e2, e3)) | ~ spl66_50), inference(backward_demodulation, [], [f254, f800])).
fof(f254, plain, ~ (op(e0, e3) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f3422, plain, (~ spl66_11 | ~ spl66_59), inference(avatar_split_clause, [], [f3419, f836, f632])).
fof(f836, plain, (spl66_59 <=> (e2 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_59])])).
fof(f3419, plain, (~ (e2 = op(e3, e1)) | ~ spl66_59), inference(backward_demodulation, [], [f243, f838])).
fof(f838, plain, ((e2 = op(e0, e1)) | ~ spl66_59), inference(avatar_component_clause, [], [f836])).
fof(f3416, plain, (~ spl66_61 | ~ spl66_133), inference(avatar_contradiction_clause, [], [f3415])).
fof(f3415, plain, ($false | (~ spl66_61 | ~ spl66_133)), inference(subsumption_resolution, [], [f3414, f285])).
fof(f3414, plain, ((e0 = e3) | (~ spl66_61 | ~ spl66_133)), inference(forward_demodulation, [], [f3410, f847])).
fof(f3410, plain, ((op(e0, e0) = e3) | (~ spl66_61 | ~ spl66_133)), inference(backward_demodulation, [], [f1440, f847])).
fof(f1440, plain, ((e3 = op(e0, op(e0, e0))) | ~ spl66_133), inference(avatar_component_clause, [], [f1439])).
fof(f1439, plain, (spl66_133 <=> (e3 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_133])])).
fof(f3411, plain, (~ spl66_29 | ~ spl66_61), inference(avatar_split_clause, [], [f3402, f845, f709])).
fof(f709, plain, (spl66_29 <=> (e0 = op(e2, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_29])])).
fof(f3402, plain, (~ (e0 = op(e2, e0)) | ~ spl66_61), inference(backward_demodulation, [], [f236, f847])).
fof(f3394, plain, (~ spl66_24 | ~ spl66_56), inference(avatar_split_clause, [], [f3393, f823, f687])).
fof(f3393, plain, (~ (e3 = op(e2, e2)) | ~ spl66_56), inference(forward_demodulation, [], [f248, f825])).
fof(f825, plain, ((e3 = op(e0, e2)) | ~ spl66_56), inference(avatar_component_clause, [], [f823])).
fof(f3385, plain, (~ spl66_6 | ~ spl66_14), inference(avatar_split_clause, [], [f3382, f645, f611])).
fof(f3382, plain, (~ (e1 = op(e3, e2)) | ~ spl66_14), inference(backward_demodulation, [], [f278, f647])).
fof(f647, plain, ((e1 = op(e3, e0)) | ~ spl66_14), inference(avatar_component_clause, [], [f645])).
fof(f3376, plain, (~ spl66_17 | ~ spl66_29), inference(avatar_split_clause, [], [f3372, f709, f658])).
fof(f658, plain, (spl66_17 <=> (e0 = op(e2, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_17])])).
fof(f3372, plain, (~ (e0 = op(e2, e3)) | ~ spl66_29), inference(backward_demodulation, [], [f273, f711])).
fof(f711, plain, ((e0 = op(e2, e0)) | ~ spl66_29), inference(avatar_component_clause, [], [f709])).
fof(f3367, plain, (~ spl66_34 | ~ spl66_38), inference(avatar_split_clause, [], [f3363, f747, f730])).
fof(f730, plain, (spl66_34 <=> (e1 = op(e1, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_34])])).
fof(f3363, plain, (~ (e1 = op(e1, e3)) | ~ spl66_38), inference(backward_demodulation, [], [f270, f749])).
fof(f749, plain, ((e1 = op(e1, e2)) | ~ spl66_38), inference(avatar_component_clause, [], [f747])).
fof(f3366, plain, (~ spl66_6 | ~ spl66_38), inference(avatar_split_clause, [], [f3362, f747, f611])).
fof(f3362, plain, (~ (e1 = op(e3, e2)) | ~ spl66_38), inference(backward_demodulation, [], [f251, f749])).
fof(f3358, plain, (~ spl66_9 | ~ spl66_41), inference(avatar_split_clause, [], [f3349, f760, f624])).
fof(f760, plain, (spl66_41 <=> (e0 = op(e1, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_41])])).
fof(f3349, plain, (~ (e0 = op(e3, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f245, f762])).
fof(f762, plain, ((e0 = op(e1, e1)) | ~ spl66_41), inference(avatar_component_clause, [], [f760])).
fof(f3347, plain, (~ spl66_40 | ~ spl66_48), inference(avatar_split_clause, [], [f3344, f789, f755])).
fof(f789, plain, (spl66_48 <=> (e3 = op(e1, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_48])])).
fof(f3344, plain, (~ (e3 = op(e1, e2)) | ~ spl66_48), inference(backward_demodulation, [], [f266, f791])).
fof(f791, plain, ((e3 = op(e1, e0)) | ~ spl66_48), inference(avatar_component_clause, [], [f789])).
fof(f3340, plain, (~ spl66_17 | ~ spl66_49), inference(avatar_split_clause, [], [f3336, f794, f658])).
fof(f3336, plain, (~ (e0 = op(e2, e3)) | ~ spl66_49), inference(backward_demodulation, [], [f254, f796])).
fof(f3330, plain, (~ spl66_26 | ~ spl66_58), inference(avatar_split_clause, [], [f3324, f832, f696])).
fof(f696, plain, (spl66_26 <=> (e1 = op(e2, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_26])])).
fof(f832, plain, (spl66_58 <=> (e1 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_58])])).
fof(f3324, plain, (~ (e1 = op(e2, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f242, f834])).
fof(f834, plain, ((e1 = op(e0, e1)) | ~ spl66_58), inference(avatar_component_clause, [], [f832])).
fof(f3322, plain, (spl66_56 | ~ spl66_63 | ~ spl66_133), inference(avatar_split_clause, [], [f3318, f1439, f853, f823])).
fof(f3318, plain, ((e3 = op(e0, e2)) | (~ spl66_63 | ~ spl66_133)), inference(backward_demodulation, [], [f1440, f855])).
fof(f3321, plain, (~ spl66_51 | ~ spl66_63), inference(avatar_split_clause, [], [f3315, f853, f802])).
fof(f3315, plain, (~ (e2 = op(e0, e3)) | ~ spl66_63), inference(backward_demodulation, [], [f261, f855])).
fof(f3320, plain, (~ spl66_15 | ~ spl66_63), inference(avatar_split_clause, [], [f3312, f853, f649])).
fof(f3312, plain, (~ (e2 = op(e3, e0)) | ~ spl66_63), inference(backward_demodulation, [], [f237, f855])).
fof(f237, plain, ~ (op(e0, e0) = op(e3, e0)), inference(cnf_transformation, [], [f3])).
fof(f3309, plain, (~ spl66_14 | ~ spl66_133 | spl66_136), inference(avatar_split_clause, [], [f3308, f1453, f1439, f645])).
fof(f1453, plain, (spl66_136 <=> (e1 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_136])])).
fof(f3308, plain, (~ (e1 = op(e3, e0)) | (~ spl66_133 | spl66_136)), inference(forward_demodulation, [], [f1455, f1440])).
fof(f1455, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | spl66_136), inference(avatar_component_clause, [], [f1453])).
fof(f3292, plain, (~ spl66_7 | ~ spl66_23), inference(avatar_split_clause, [], [f3291, f683, f615])).
fof(f3291, plain, (~ (e2 = op(e3, e2)) | ~ spl66_23), inference(forward_demodulation, [], [f252, f685])).
fof(f685, plain, ((e2 = op(e2, e2)) | ~ spl66_23), inference(avatar_component_clause, [], [f683])).
fof(f3277, plain, (~ spl66_19 | ~ spl66_23), inference(avatar_split_clause, [], [f3272, f683, f666])).
fof(f3272, plain, (~ (e2 = op(e2, e3)) | ~ spl66_23), inference(backward_demodulation, [], [f276, f685])).
fof(f3254, plain, (~ spl66_11 | ~ spl66_43), inference(avatar_split_clause, [], [f3249, f768, f632])).
fof(f3249, plain, (~ (e2 = op(e3, e1)) | ~ spl66_43), inference(backward_demodulation, [], [f245, f770])).
fof(f3246, plain, (~ spl66_13 | ~ spl66_45), inference(avatar_split_clause, [], [f3240, f777, f641])).
fof(f641, plain, (spl66_13 <=> (e0 = op(e3, e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_13])])).
fof(f3240, plain, (~ (e0 = op(e3, e0)) | ~ spl66_45), inference(backward_demodulation, [], [f239, f779])).
fof(f779, plain, ((e0 = op(e1, e0)) | ~ spl66_45), inference(avatar_component_clause, [], [f777])).
fof(f3238, plain, (~ spl66_19 | ~ spl66_51), inference(avatar_split_clause, [], [f3236, f802, f666])).
fof(f3236, plain, (~ (e2 = op(e2, e3)) | ~ spl66_51), inference(backward_demodulation, [], [f254, f804])).
fof(f804, plain, ((e2 = op(e0, e3)) | ~ spl66_51), inference(avatar_component_clause, [], [f802])).
fof(f3234, plain, (~ spl66_21 | ~ spl66_53), inference(avatar_split_clause, [], [f3229, f811, f675])).
fof(f811, plain, (spl66_53 <=> (e0 = op(e0, e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_53])])).
fof(f3229, plain, (~ (e0 = op(e2, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f248, f813])).
fof(f813, plain, ((e0 = op(e0, e2)) | ~ spl66_53), inference(avatar_component_clause, [], [f811])).
fof(f3222, plain, (spl66_60 | ~ spl66_62 | ~ spl66_133), inference(avatar_split_clause, [], [f3219, f1439, f849, f840])).
fof(f849, plain, (spl66_62 <=> (op(e0, e0) = e1)), introduced(avatar_definition, [new_symbols(naming, [spl66_62])])).
fof(f3219, plain, ((e3 = op(e0, e1)) | (~ spl66_62 | ~ spl66_133)), inference(backward_demodulation, [], [f1440, f851])).
fof(f851, plain, ((op(e0, e0) = e1) | ~ spl66_62), inference(avatar_component_clause, [], [f849])).
fof(f3221, plain, (~ spl66_50 | ~ spl66_62), inference(avatar_split_clause, [], [f3215, f849, f798])).
fof(f3215, plain, (~ (e1 = op(e0, e3)) | ~ spl66_62), inference(backward_demodulation, [], [f261, f851])).
fof(f3220, plain, (~ spl66_30 | ~ spl66_62), inference(avatar_split_clause, [], [f3211, f849, f713])).
fof(f3211, plain, (~ (e1 = op(e2, e0)) | ~ spl66_62), inference(backward_demodulation, [], [f236, f851])).
fof(f3209, plain, (~ spl66_15 | ~ spl66_133 | spl66_134), inference(avatar_split_clause, [], [f3208, f1443, f1439, f649])).
fof(f1443, plain, (spl66_134 <=> (e2 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_134])])).
fof(f3208, plain, (~ (e2 = op(e3, e0)) | (~ spl66_133 | spl66_134)), inference(backward_demodulation, [], [f1445, f1440])).
fof(f1445, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | spl66_134), inference(avatar_component_clause, [], [f1443])).
fof(f3190, plain, (~ spl66_5 | ~ spl66_13), inference(avatar_split_clause, [], [f3187, f641, f607])).
fof(f3187, plain, (~ (e0 = op(e3, e2)) | ~ spl66_13), inference(backward_demodulation, [], [f278, f643])).
fof(f643, plain, ((e0 = op(e3, e0)) | ~ spl66_13), inference(avatar_component_clause, [], [f641])).
fof(f3160, plain, (~ spl66_56 | ~ spl66_64), inference(avatar_split_clause, [], [f3157, f857, f823])).
fof(f857, plain, (spl66_64 <=> (op(e0, e0) = e3)), introduced(avatar_definition, [new_symbols(naming, [spl66_64])])).
fof(f3157, plain, (~ (e3 = op(e0, e2)) | ~ spl66_64), inference(backward_demodulation, [], [f260, f859])).
fof(f859, plain, ((op(e0, e0) = e3) | ~ spl66_64), inference(avatar_component_clause, [], [f857])).
fof(f3153, plain, (~ spl66_47 | spl66_134 | ~ spl66_135), inference(avatar_split_clause, [], [f3152, f1448, f1443, f785])).
fof(f1448, plain, (spl66_135 <=> (e1 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_135])])).
fof(f3152, plain, (~ (e2 = op(e1, e0)) | (spl66_134 | ~ spl66_135)), inference(forward_demodulation, [], [f1445, f1449])).
fof(f1449, plain, ((e1 = op(e0, op(e0, e0))) | ~ spl66_135), inference(avatar_component_clause, [], [f1448])).
fof(f3151, plain, (~ spl66_48 | spl66_132 | ~ spl66_135), inference(avatar_split_clause, [], [f3149, f1448, f1434, f789])).
fof(f1434, plain, (spl66_132 <=> (e3 = op(op(e0, op(e0, e0)), e0))), introduced(avatar_definition, [new_symbols(naming, [spl66_132])])).
fof(f3149, plain, (~ (e3 = op(e1, e0)) | (spl66_132 | ~ spl66_135)), inference(backward_demodulation, [], [f1436, f1449])).
fof(f1436, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | spl66_132), inference(avatar_component_clause, [], [f1434])).
fof(f3116, plain, (~ spl66_24 | ~ spl66_28), inference(avatar_split_clause, [], [f3114, f704, f687])).
fof(f3114, plain, (~ (e3 = op(e2, e2)) | ~ spl66_28), inference(backward_demodulation, [], [f274, f706])).
fof(f706, plain, ((e3 = op(e2, e1)) | ~ spl66_28), inference(avatar_component_clause, [], [f704])).
fof(f274, plain, ~ (op(e2, e1) = op(e2, e2)), inference(cnf_transformation, [], [f3])).
fof(f3106, plain, (~ spl66_7 | ~ spl66_39), inference(avatar_split_clause, [], [f3104, f751, f615])).
fof(f3104, plain, (~ (e2 = op(e3, e2)) | ~ spl66_39), inference(backward_demodulation, [], [f251, f753])).
fof(f3089, plain, (~ spl66_49 | ~ spl66_57), inference(avatar_split_clause, [], [f3084, f828, f794])).
fof(f828, plain, (spl66_57 <=> (e0 = op(e0, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_57])])).
fof(f3084, plain, (~ (e0 = op(e0, e3)) | ~ spl66_57), inference(backward_demodulation, [], [f263, f830])).
fof(f830, plain, ((e0 = op(e0, e1)) | ~ spl66_57), inference(avatar_component_clause, [], [f828])).
fof(f3088, plain, (~ spl66_9 | ~ spl66_57), inference(avatar_split_clause, [], [f3082, f828, f624])).
fof(f3082, plain, (~ (e0 = op(e3, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f243, f830])).
fof(f3080, plain, (~ spl66_5 | ~ spl66_145 | spl66_148), inference(avatar_split_clause, [], [f3079, f1513, f1499, f607])).
fof(f1499, plain, (spl66_145 <=> (e3 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_145])])).
fof(f1513, plain, (spl66_148 <=> (e0 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_148])])).
fof(f3079, plain, (~ (e0 = op(e3, e2)) | (~ spl66_145 | spl66_148)), inference(backward_demodulation, [], [f1515, f1500])).
fof(f1500, plain, ((e3 = op(e2, op(e2, e2))) | ~ spl66_145), inference(avatar_component_clause, [], [f1499])).
fof(f1515, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | spl66_148), inference(avatar_component_clause, [], [f1513])).
fof(f3077, plain, (~ spl66_56 | ~ spl66_63 | spl66_133), inference(avatar_split_clause, [], [f3021, f1439, f853, f823])).
fof(f3021, plain, (~ (e3 = op(e0, e2)) | (~ spl66_63 | spl66_133)), inference(backward_demodulation, [], [f1441, f855])).
fof(f1441, plain, (~ (e3 = op(e0, op(e0, e0))) | spl66_133), inference(avatar_component_clause, [], [f1439])).
fof(f3076, plain, (~ spl66_52 | ~ spl66_4), inference(avatar_split_clause, [], [f2926, f602, f806])).
fof(f2926, plain, (~ (e3 = op(e0, e3)) | ~ spl66_4), inference(forward_demodulation, [], [f255, f604])).
fof(f604, plain, ((e3 = op(e3, e3)) | ~ spl66_4), inference(avatar_component_clause, [], [f602])).
fof(f3073, plain, (~ spl66_36 | ~ spl66_4), inference(avatar_split_clause, [], [f2928, f602, f738])).
fof(f2928, plain, (~ (e3 = op(e1, e3)) | ~ spl66_4), inference(forward_demodulation, [], [f257, f604])).
fof(f3068, plain, (~ spl66_12 | ~ spl66_4), inference(avatar_split_clause, [], [f2932, f602, f636])).
fof(f2932, plain, (~ (e3 = op(e3, e1)) | ~ spl66_4), inference(forward_demodulation, [], [f281, f604])).
fof(f281, plain, ~ (op(e3, e1) = op(e3, e3)), inference(cnf_transformation, [], [f3])).
fof(f3067, plain, (~ spl66_10 | ~ spl66_14), inference(avatar_split_clause, [], [f3066, f645, f628])).
fof(f628, plain, (spl66_10 <=> (e1 = op(e3, e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_10])])).
fof(f3066, plain, (~ (e1 = op(e3, e1)) | ~ spl66_14), inference(forward_demodulation, [], [f277, f647])).
fof(f3060, plain, (~ spl66_18 | ~ spl66_24 | ~ spl66_37 | spl66_148), inference(avatar_contradiction_clause, [], [f3059])).
fof(f3059, plain, ($false | (~ spl66_18 | ~ spl66_24 | ~ spl66_37 | spl66_148)), inference(subsumption_resolution, [], [f3058, f745])).
fof(f3058, plain, (~ (e0 = op(e1, e2)) | (~ spl66_18 | ~ spl66_24 | spl66_148)), inference(forward_demodulation, [], [f3055, f664])).
fof(f664, plain, ((e1 = op(e2, e3)) | ~ spl66_18), inference(avatar_component_clause, [], [f662])).
fof(f3055, plain, (~ (e0 = op(op(e2, e3), e2)) | (~ spl66_24 | spl66_148)), inference(backward_demodulation, [], [f1515, f689])).
fof(f689, plain, ((e3 = op(e2, e2)) | ~ spl66_24), inference(avatar_component_clause, [], [f687])).
fof(f3057, plain, (~ spl66_18 | ~ spl66_24 | spl66_143), inference(avatar_contradiction_clause, [], [f3056])).
fof(f3056, plain, ($false | (~ spl66_18 | ~ spl66_24 | spl66_143)), inference(subsumption_resolution, [], [f3053, f664])).
fof(f3053, plain, (~ (e1 = op(e2, e3)) | (~ spl66_24 | spl66_143)), inference(backward_demodulation, [], [f1492, f689])).
fof(f1492, plain, (~ (e1 = op(e2, op(e2, e2))) | spl66_143), inference(avatar_component_clause, [], [f1490])).
fof(f1490, plain, (spl66_143 <=> (e1 = op(e2, op(e2, e2)))), introduced(avatar_definition, [new_symbols(naming, [spl66_143])])).
fof(f3051, plain, (~ spl66_21 | ~ spl66_29), inference(avatar_split_clause, [], [f3048, f709, f675])).
fof(f3048, plain, (~ (e0 = op(e2, e2)) | ~ spl66_29), inference(backward_demodulation, [], [f272, f711])).
fof(f3027, plain, (~ spl66_54 | ~ spl66_63 | spl66_135), inference(avatar_split_clause, [], [f3022, f1448, f853, f815])).
fof(f3022, plain, (~ (e1 = op(e0, e2)) | (~ spl66_63 | spl66_135)), inference(backward_demodulation, [], [f1450, f855])).
fof(f1450, plain, (~ (e1 = op(e0, op(e0, e0))) | spl66_135), inference(avatar_component_clause, [], [f1448])).
fof(f3005, plain, (~ spl66_31 | ~ spl66_27), inference(avatar_split_clause, [], [f2960, f700, f717])).
fof(f2960, plain, (~ (e2 = op(e2, e0)) | ~ spl66_27), inference(forward_demodulation, [], [f271, f702])).
fof(f702, plain, ((e2 = op(e2, e1)) | ~ spl66_27), inference(avatar_component_clause, [], [f700])).
fof(f3003, plain, (~ spl66_23 | ~ spl66_27), inference(avatar_split_clause, [], [f3002, f700, f683])).
fof(f3002, plain, (~ (e2 = op(e2, e2)) | ~ spl66_27), inference(forward_demodulation, [], [f274, f702])).
fof(f2993, plain, (~ spl66_6 | ~ spl66_21 | ~ spl66_32 | spl66_146), inference(avatar_split_clause, [], [f2990, f1503, f721, f675, f611])).
fof(f1503, plain, (spl66_146 <=> (e1 = op(op(e2, op(e2, e2)), e2))), introduced(avatar_definition, [new_symbols(naming, [spl66_146])])).
fof(f2990, plain, (~ (e1 = op(e3, e2)) | (~ spl66_21 | ~ spl66_32 | spl66_146)), inference(backward_demodulation, [], [f2892, f723])).
fof(f2892, plain, (~ (e1 = op(op(e2, e0), e2)) | (~ spl66_21 | spl66_146)), inference(backward_demodulation, [], [f1505, f677])).
fof(f1505, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | spl66_146), inference(avatar_component_clause, [], [f1503])).
fof(f2985, plain, (~ spl66_60 | ~ spl66_62 | spl66_133), inference(avatar_split_clause, [], [f2981, f1439, f849, f840])).
fof(f2981, plain, (~ (e3 = op(e0, e1)) | (~ spl66_62 | spl66_133)), inference(backward_demodulation, [], [f1441, f851])).
fof(f2984, plain, (~ spl66_58 | ~ spl66_62), inference(avatar_split_clause, [], [f2976, f849, f832])).
fof(f2976, plain, (~ (e1 = op(e0, e1)) | ~ spl66_62), inference(backward_demodulation, [], [f259, f851])).
fof(f259, plain, ~ (op(e0, e0) = op(e0, e1)), inference(cnf_transformation, [], [f3])).
fof(f2967, plain, (~ spl66_59 | ~ spl66_55), inference(avatar_split_clause, [], [f2966, f819, f836])).
fof(f2966, plain, (~ (e2 = op(e0, e1)) | ~ spl66_55), inference(forward_demodulation, [], [f262, f821])).
fof(f2955, plain, (~ spl66_30 | ~ spl66_21 | spl66_143), inference(avatar_split_clause, [], [f2890, f1490, f675, f713])).
fof(f2890, plain, (~ (e1 = op(e2, e0)) | (~ spl66_21 | spl66_143)), inference(backward_demodulation, [], [f1492, f677])).
fof(f2954, plain, (~ spl66_32 | ~ spl66_21 | spl66_145), inference(avatar_split_clause, [], [f2891, f1499, f675, f721])).
fof(f2891, plain, (~ (e3 = op(e2, e0)) | (~ spl66_21 | spl66_145)), inference(backward_demodulation, [], [f1501, f677])).
fof(f1501, plain, (~ (e3 = op(e2, op(e2, e2))) | spl66_145), inference(avatar_component_clause, [], [f1499])).
fof(f2950, plain, (~ spl66_59 | ~ spl66_27), inference(avatar_split_clause, [], [f2949, f700, f836])).
fof(f2949, plain, (~ (e2 = op(e0, e1)) | ~ spl66_27), inference(forward_demodulation, [], [f242, f702])).
fof(f2916, plain, (~ spl66_2 | ~ spl66_4), inference(avatar_contradiction_clause, [], [f2915])).
fof(f2915, plain, ($false | (~ spl66_2 | ~ spl66_4)), inference(subsumption_resolution, [], [f2914, f287])).
fof(f287, plain, ~ (e1 = e3), inference(cnf_transformation, [], [f4])).
fof(f2914, plain, ((e1 = e3) | (~ spl66_2 | ~ spl66_4)), inference(forward_demodulation, [], [f604, f596])).
fof(f596, plain, ((e1 = op(e3, e3)) | ~ spl66_2), inference(avatar_component_clause, [], [f594])).
fof(f2877, plain, (~ spl66_37 | ~ spl66_45), inference(avatar_split_clause, [], [f2872, f777, f743])).
fof(f2872, plain, (~ (e0 = op(e1, e2)) | ~ spl66_45), inference(backward_demodulation, [], [f266, f779])).
fof(f2853, plain, (~ spl66_51 | ~ spl66_64 | spl66_131), inference(avatar_split_clause, [], [f2848, f1430, f857, f802])).
fof(f1430, plain, (spl66_131 <=> (e2 = op(e0, op(e0, e0)))), introduced(avatar_definition, [new_symbols(naming, [spl66_131])])).
fof(f2848, plain, (~ (e2 = op(e0, e3)) | (~ spl66_64 | spl66_131)), inference(backward_demodulation, [], [f1432, f859])).
fof(f1432, plain, (~ (e2 = op(e0, op(e0, e0))) | spl66_131), inference(avatar_component_clause, [], [f1430])).
fof(f2852, plain, (~ spl66_16 | ~ spl66_64), inference(avatar_split_clause, [], [f2845, f857, f653])).
fof(f2845, plain, (~ (e3 = op(e3, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f237, f859])).
fof(f2838, plain, (~ spl66_58 | ~ spl66_54), inference(avatar_split_clause, [], [f2715, f815, f832])).
fof(f2715, plain, (~ (e1 = op(e0, e1)) | ~ spl66_54), inference(forward_demodulation, [], [f262, f817])).
fof(f2835, plain, (~ spl66_58 | ~ spl66_42), inference(avatar_split_clause, [], [f2834, f764, f832])).
fof(f2834, plain, (~ (e1 = op(e0, e1)) | ~ spl66_42), inference(forward_demodulation, [], [f241, f766])).
fof(f766, plain, ((e1 = op(e1, e1)) | ~ spl66_42), inference(avatar_component_clause, [], [f764])).
fof(f2824, plain, (~ spl66_18 | ~ spl66_2), inference(avatar_split_clause, [], [f2823, f594, f662])).
fof(f2823, plain, (~ (e1 = op(e2, e3)) | ~ spl66_2), inference(forward_demodulation, [], [f258, f596])).
fof(f2822, plain, (~ spl66_20 | ~ spl66_28), inference(avatar_split_clause, [], [f2821, f704, f670])).
fof(f2821, plain, (~ (e3 = op(e2, e3)) | ~ spl66_28), inference(forward_demodulation, [], [f275, f706])).
fof(f275, plain, ~ (op(e2, e1) = op(e2, e3)), inference(cnf_transformation, [], [f3])).
fof(f2820, plain, (~ spl66_18 | ~ spl66_30), inference(avatar_split_clause, [], [f2819, f713, f662])).
fof(f2819, plain, (~ (e1 = op(e2, e3)) | ~ spl66_30), inference(forward_demodulation, [], [f273, f715])).
fof(f2817, plain, (~ spl66_12 | ~ spl66_28), inference(avatar_split_clause, [], [f2816, f704, f636])).
fof(f2816, plain, (~ (e3 = op(e3, e1)) | ~ spl66_28), inference(forward_demodulation, [], [f246, f706])).
fof(f246, plain, ~ (op(e2, e1) = op(e3, e1)), inference(cnf_transformation, [], [f3])).
fof(f2813, plain, (~ spl66_9 | ~ spl66_2 | spl66_153), inference(avatar_split_clause, [], [f2673, f1538, f594, f624])).
fof(f2673, plain, (~ (e0 = op(e3, e1)) | (~ spl66_2 | spl66_153)), inference(backward_demodulation, [], [f1540, f596])).
fof(f2811, plain, (~ spl66_26 | ~ spl66_30), inference(avatar_split_clause, [], [f2809, f713, f696])).
fof(f2809, plain, (~ (e1 = op(e2, e1)) | ~ spl66_30), inference(backward_demodulation, [], [f271, f715])).
fof(f2799, plain, (~ spl66_38 | ~ spl66_42), inference(avatar_split_clause, [], [f2792, f764, f747])).
fof(f2792, plain, (~ (e1 = op(e1, e2)) | ~ spl66_42), inference(backward_demodulation, [], [f268, f766])).
fof(f2798, plain, (~ spl66_26 | ~ spl66_42), inference(avatar_split_clause, [], [f2791, f764, f696])).
fof(f2791, plain, (~ (e1 = op(e2, e1)) | ~ spl66_42), inference(backward_demodulation, [], [f244, f766])).
fof(f2787, plain, (~ spl66_41 | ~ spl66_57), inference(avatar_split_clause, [], [f2784, f828, f760])).
fof(f2784, plain, (~ (e0 = op(e1, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f241, f830])).
fof(f2774, plain, (~ spl66_57 | ~ spl66_61), inference(avatar_split_clause, [], [f2713, f845, f828])).
fof(f2713, plain, (~ (e0 = op(e0, e1)) | ~ spl66_61), inference(forward_demodulation, [], [f259, f847])).
fof(f2772, plain, (~ spl66_43 | ~ spl66_47), inference(avatar_split_clause, [], [f2771, f785, f768])).
fof(f2771, plain, (~ (e2 = op(e1, e1)) | ~ spl66_47), inference(forward_demodulation, [], [f265, f787])).
fof(f2767, plain, (~ spl66_38 | ~ spl66_54), inference(avatar_split_clause, [], [f2766, f815, f747])).
fof(f2766, plain, (~ (e1 = op(e1, e2)) | ~ spl66_54), inference(forward_demodulation, [], [f247, f817])).
fof(f2761, plain, (~ spl66_44 | ~ spl66_36), inference(avatar_split_clause, [], [f2760, f738, f772])).
fof(f2760, plain, (~ (e3 = op(e1, e1)) | ~ spl66_36), inference(forward_demodulation, [], [f269, f740])).
fof(f269, plain, ~ (op(e1, e1) = op(e1, e3)), inference(cnf_transformation, [], [f3])).
fof(f2757, plain, (~ spl66_34 | ~ spl66_2), inference(avatar_split_clause, [], [f2704, f594, f730])).
fof(f2704, plain, (~ (e1 = op(e1, e3)) | ~ spl66_2), inference(forward_demodulation, [], [f257, f596])).
fof(f2747, plain, (~ spl66_14 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f2746])).
fof(f2746, plain, ($false | (~ spl66_14 | ~ spl66_16)), inference(subsumption_resolution, [], [f2745, f287])).
fof(f2745, plain, ((e1 = e3) | (~ spl66_14 | ~ spl66_16)), inference(backward_demodulation, [], [f655, f647])).
fof(f2743, plain, (~ spl66_10 | ~ spl66_26), inference(avatar_split_clause, [], [f2741, f696, f628])).
fof(f2741, plain, (~ (e1 = op(e3, e1)) | ~ spl66_26), inference(backward_demodulation, [], [f246, f698])).
fof(f698, plain, ((e1 = op(e2, e1)) | ~ spl66_26), inference(avatar_component_clause, [], [f696])).
fof(f2736, plain, (~ spl66_34 | ~ spl66_36), inference(avatar_contradiction_clause, [], [f2735])).
fof(f2735, plain, ($false | (~ spl66_34 | ~ spl66_36)), inference(subsumption_resolution, [], [f2734, f287])).
fof(f2734, plain, ((e1 = e3) | (~ spl66_34 | ~ spl66_36)), inference(forward_demodulation, [], [f740, f732])).
fof(f732, plain, ((e1 = op(e1, e3)) | ~ spl66_34), inference(avatar_component_clause, [], [f730])).
fof(f2733, plain, (~ spl66_28 | ~ spl66_41 | ~ spl66_47 | spl66_138), inference(avatar_split_clause, [], [f2732, f1464, f785, f760, f704])).
fof(f1464, plain, (spl66_138 <=> (e3 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_138])])).
fof(f2732, plain, (~ (e3 = op(e2, e1)) | (~ spl66_41 | ~ spl66_47 | spl66_138)), inference(backward_demodulation, [], [f2716, f787])).
fof(f2716, plain, (~ (e3 = op(op(e1, e0), e1)) | (~ spl66_41 | spl66_138)), inference(forward_demodulation, [], [f1466, f762])).
fof(f1466, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | spl66_138), inference(avatar_component_clause, [], [f1464])).
fof(f2694, plain, (~ spl66_14 | ~ spl66_2), inference(avatar_split_clause, [], [f2693, f594, f645])).
fof(f2693, plain, (~ (e1 = op(e3, e0)) | ~ spl66_2), inference(forward_demodulation, [], [f279, f596])).
fof(f2690, plain, (~ spl66_48 | ~ spl66_16), inference(avatar_split_clause, [], [f2689, f653, f789])).
fof(f2689, plain, (~ (e3 = op(e1, e0)) | ~ spl66_16), inference(forward_demodulation, [], [f239, f655])).
fof(f2685, plain, (~ spl66_10 | ~ spl66_2), inference(avatar_split_clause, [], [f2684, f594, f628])).
fof(f2684, plain, (~ (e1 = op(e3, e1)) | ~ spl66_2), inference(forward_demodulation, [], [f281, f596])).
fof(f2667, plain, (~ spl66_1 | ~ spl66_5), inference(avatar_split_clause, [], [f2664, f607, f590])).
fof(f590, plain, (spl66_1 <=> (e0 = op(e3, e3))), introduced(avatar_definition, [new_symbols(naming, [spl66_1])])).
fof(f2664, plain, (~ (e0 = op(e3, e3)) | ~ spl66_5), inference(backward_demodulation, [], [f282, f609])).
fof(f2661, plain, (~ spl66_6 | ~ spl66_10), inference(avatar_split_clause, [], [f2657, f628, f611])).
fof(f2657, plain, (~ (e1 = op(e3, e2)) | ~ spl66_10), inference(backward_demodulation, [], [f280, f630])).
fof(f630, plain, ((e1 = op(e3, e1)) | ~ spl66_10), inference(avatar_component_clause, [], [f628])).
fof(f280, plain, ~ (op(e3, e1) = op(e3, e2)), inference(cnf_transformation, [], [f3])).
fof(f2656, plain, (~ spl66_15 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f2655])).
fof(f2655, plain, ($false | (~ spl66_15 | ~ spl66_16)), inference(subsumption_resolution, [], [f2654, f288])).
fof(f288, plain, ~ (e2 = e3), inference(cnf_transformation, [], [f4])).
fof(f2654, plain, ((e2 = e3) | (~ spl66_15 | ~ spl66_16)), inference(forward_demodulation, [], [f655, f651])).
fof(f2653, plain, (~ spl66_1 | ~ spl66_17), inference(avatar_split_clause, [], [f2652, f658, f590])).
fof(f2652, plain, (~ (e0 = op(e3, e3)) | ~ spl66_17), inference(backward_demodulation, [], [f258, f660])).
fof(f660, plain, ((e0 = op(e2, e3)) | ~ spl66_17), inference(avatar_component_clause, [], [f658])).
fof(f2640, plain, (~ spl66_34 | ~ spl66_35), inference(avatar_contradiction_clause, [], [f2639])).
fof(f2639, plain, ($false | (~ spl66_34 | ~ spl66_35)), inference(subsumption_resolution, [], [f2638, f286])).
fof(f286, plain, ~ (e1 = e2), inference(cnf_transformation, [], [f4])).
fof(f2638, plain, ((e1 = e2) | (~ spl66_34 | ~ spl66_35)), inference(backward_demodulation, [], [f736, f732])).
fof(f2635, plain, (~ spl66_25 | ~ spl66_41), inference(avatar_split_clause, [], [f2627, f760, f692])).
fof(f2627, plain, (~ (e0 = op(e2, e1)) | ~ spl66_41), inference(backward_demodulation, [], [f244, f762])).
fof(f2568, plain, (~ spl66_20 | ~ spl66_52), inference(avatar_split_clause, [], [f2567, f806, f670])).
fof(f2567, plain, (~ (e3 = op(e2, e3)) | ~ spl66_52), inference(backward_demodulation, [], [f254, f808])).
fof(f808, plain, ((e3 = op(e0, e3)) | ~ spl66_52), inference(avatar_component_clause, [], [f806])).
fof(f2566, plain, (~ spl66_37 | ~ spl66_53), inference(avatar_split_clause, [], [f2561, f811, f743])).
fof(f2561, plain, (~ (e0 = op(e1, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f247, f813])).
fof(f2560, plain, (~ spl66_58 | ~ spl66_59), inference(avatar_contradiction_clause, [], [f2559])).
fof(f2559, plain, ($false | (~ spl66_58 | ~ spl66_59)), inference(subsumption_resolution, [], [f2558, f286])).
fof(f2558, plain, ((e1 = e2) | (~ spl66_58 | ~ spl66_59)), inference(backward_demodulation, [], [f838, f834])).
fof(f2556, plain, (~ spl66_61 | ~ spl66_131), inference(avatar_contradiction_clause, [], [f2555])).
fof(f2555, plain, ($false | (~ spl66_61 | ~ spl66_131)), inference(subsumption_resolution, [], [f2554, f284])).
fof(f284, plain, ~ (e0 = e2), inference(cnf_transformation, [], [f4])).
fof(f2554, plain, ((e0 = e2) | (~ spl66_61 | ~ spl66_131)), inference(forward_demodulation, [], [f2550, f847])).
fof(f2550, plain, ((op(e0, e0) = e2) | (~ spl66_61 | ~ spl66_131)), inference(backward_demodulation, [], [f1431, f847])).
fof(f1431, plain, ((e2 = op(e0, op(e0, e0))) | ~ spl66_131), inference(avatar_component_clause, [], [f1430])).
fof(f2553, plain, (~ spl66_53 | ~ spl66_61), inference(avatar_split_clause, [], [f2545, f845, f811])).
fof(f2545, plain, (~ (e0 = op(e0, e2)) | ~ spl66_61), inference(backward_demodulation, [], [f260, f847])).
fof(f2539, plain, (~ spl66_62 | ~ spl66_46), inference(avatar_split_clause, [], [f2538, f781, f849])).
fof(f2538, plain, (~ (op(e0, e0) = e1) | ~ spl66_46), inference(forward_demodulation, [], [f235, f783])).
fof(f2493, plain, (~ spl66_8 | ~ spl66_12), inference(avatar_split_clause, [], [f2491, f636, f619])).
fof(f2491, plain, (~ (e3 = op(e3, e2)) | ~ spl66_12), inference(backward_demodulation, [], [f280, f638])).
fof(f2479, plain, (~ spl66_25 | ~ spl66_29), inference(avatar_split_clause, [], [f2475, f709, f692])).
fof(f2475, plain, (~ (e0 = op(e2, e1)) | ~ spl66_29), inference(backward_demodulation, [], [f271, f711])).
fof(f2472, plain, (~ spl66_25 | ~ spl66_35 | ~ spl66_44 | spl66_142), inference(avatar_split_clause, [], [f2470, f1483, f772, f734, f692])).
fof(f1483, plain, (spl66_142 <=> (e0 = op(op(e1, op(e1, e1)), e1))), introduced(avatar_definition, [new_symbols(naming, [spl66_142])])).
fof(f2470, plain, (~ (e0 = op(e2, e1)) | (~ spl66_35 | ~ spl66_44 | spl66_142)), inference(backward_demodulation, [], [f2218, f736])).
fof(f2218, plain, (~ (e0 = op(op(e1, e3), e1)) | (~ spl66_44 | spl66_142)), inference(backward_demodulation, [], [f1485, f774])).
fof(f1485, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | spl66_142), inference(avatar_component_clause, [], [f1483])).
fof(f2459, plain, (~ spl66_46 | ~ spl66_48), inference(avatar_contradiction_clause, [], [f2458])).
fof(f2458, plain, ($false | (~ spl66_46 | ~ spl66_48)), inference(subsumption_resolution, [], [f2457, f287])).
fof(f2457, plain, ((e1 = e3) | (~ spl66_46 | ~ spl66_48)), inference(backward_demodulation, [], [f791, f783])).
fof(f2443, plain, (~ spl66_8 | ~ spl66_56), inference(avatar_split_clause, [], [f2441, f823, f619])).
fof(f2441, plain, (~ (e3 = op(e3, e2)) | ~ spl66_56), inference(backward_demodulation, [], [f249, f825])).
fof(f2439, plain, (~ spl66_57 | ~ spl66_59), inference(avatar_contradiction_clause, [], [f2438])).
fof(f2438, plain, ($false | (~ spl66_57 | ~ spl66_59)), inference(subsumption_resolution, [], [f2437, f284])).
fof(f2437, plain, ((e0 = e2) | (~ spl66_57 | ~ spl66_59)), inference(backward_demodulation, [], [f838, f830])).
fof(f2412, plain, (~ spl66_32 | ~ spl66_131 | spl66_132), inference(avatar_split_clause, [], [f2409, f1434, f1430, f721])).
fof(f2409, plain, (~ (e3 = op(e2, e0)) | (~ spl66_131 | spl66_132)), inference(backward_demodulation, [], [f1436, f1431])).
fof(f2408, plain, (~ spl66_48 | ~ spl66_44), inference(avatar_split_clause, [], [f2407, f772, f789])).
fof(f2407, plain, (~ (e3 = op(e1, e0)) | ~ spl66_44), inference(forward_demodulation, [], [f265, f774])).
fof(f2395, plain, (~ spl66_3 | ~ spl66_11), inference(avatar_split_clause, [], [f2393, f632, f598])).
fof(f2393, plain, (~ (e2 = op(e3, e3)) | ~ spl66_11), inference(backward_demodulation, [], [f281, f634])).
fof(f634, plain, ((e2 = op(e3, e1)) | ~ spl66_11), inference(avatar_component_clause, [], [f632])).
fof(f2391, plain, (~ spl66_3 | ~ spl66_19), inference(avatar_split_clause, [], [f2390, f666, f598])).
fof(f2390, plain, (~ (e2 = op(e3, e3)) | ~ spl66_19), inference(backward_demodulation, [], [f258, f668])).
fof(f668, plain, ((e2 = op(e2, e3)) | ~ spl66_19), inference(avatar_component_clause, [], [f666])).
fof(f2388, plain, (~ spl66_17 | ~ spl66_25), inference(avatar_split_clause, [], [f2386, f692, f658])).
fof(f2386, plain, (~ (e0 = op(e2, e3)) | ~ spl66_25), inference(backward_demodulation, [], [f275, f694])).
fof(f2382, plain, (~ spl66_30 | ~ spl66_32), inference(avatar_contradiction_clause, [], [f2381])).
fof(f2381, plain, ($false | (~ spl66_30 | ~ spl66_32)), inference(subsumption_resolution, [], [f2380, f287])).
fof(f2380, plain, ((e1 = e3) | (~ spl66_30 | ~ spl66_32)), inference(backward_demodulation, [], [f723, f715])).
fof(f2369, plain, (~ spl66_10 | ~ spl66_58), inference(avatar_split_clause, [], [f2367, f832, f628])).
fof(f2367, plain, (~ (e1 = op(e3, e1)) | ~ spl66_58), inference(backward_demodulation, [], [f243, f834])).
fof(f2356, plain, (~ spl66_59 | ~ spl66_63), inference(avatar_split_clause, [], [f2349, f853, f836])).
fof(f2349, plain, (~ (e2 = op(e0, e1)) | ~ spl66_63), inference(backward_demodulation, [], [f259, f855])).
fof(f2345, plain, (~ spl66_64 | ~ spl66_52), inference(avatar_split_clause, [], [f2271, f806, f857])).
fof(f2271, plain, (~ (op(e0, e0) = e3) | ~ spl66_52), inference(forward_demodulation, [], [f261, f808])).
fof(f2344, plain, (~ spl66_63 | ~ spl66_47), inference(avatar_split_clause, [], [f2178, f785, f853])).
fof(f2178, plain, (~ (op(e0, e0) = e2) | ~ spl66_47), inference(forward_demodulation, [], [f235, f787])).
fof(f2339, plain, (~ spl66_13 | ~ spl66_16), inference(avatar_contradiction_clause, [], [f2338])).
fof(f2338, plain, ($false | (~ spl66_13 | ~ spl66_16)), inference(subsumption_resolution, [], [f2337, f285])).
fof(f2337, plain, ((e0 = e3) | (~ spl66_13 | ~ spl66_16)), inference(forward_demodulation, [], [f655, f643])).
fof(f2333, plain, (~ spl66_18 | ~ spl66_26), inference(avatar_split_clause, [], [f2332, f696, f662])).
fof(f2332, plain, (~ (e1 = op(e2, e3)) | ~ spl66_26), inference(backward_demodulation, [], [f275, f698])).
fof(f2314, plain, (~ spl66_50 | ~ spl66_52), inference(avatar_contradiction_clause, [], [f2313])).
fof(f2313, plain, ($false | (~ spl66_50 | ~ spl66_52)), inference(subsumption_resolution, [], [f2312, f287])).
fof(f2312, plain, ((e1 = e3) | (~ spl66_50 | ~ spl66_52)), inference(backward_demodulation, [], [f808, f800])).
fof(f2298, plain, (~ spl66_59 | ~ spl66_62 | spl66_131), inference(avatar_split_clause, [], [f2297, f1430, f849, f836])).
fof(f2297, plain, (~ (e2 = op(e0, e1)) | (~ spl66_62 | spl66_131)), inference(forward_demodulation, [], [f1432, f851])).
fof(f2274, plain, (~ spl66_36 | ~ spl66_52), inference(avatar_split_clause, [], [f2273, f806, f738])).
fof(f2273, plain, (~ (e3 = op(e1, e3)) | ~ spl66_52), inference(forward_demodulation, [], [f253, f808])).
fof(f2248, plain, (~ spl66_4 | ~ spl66_8), inference(avatar_split_clause, [], [f2247, f619, f602])).
fof(f2247, plain, (~ (e3 = op(e3, e3)) | ~ spl66_8), inference(backward_demodulation, [], [f282, f621])).
fof(f621, plain, ((e3 = op(e3, e2)) | ~ spl66_8), inference(avatar_component_clause, [], [f619])).
fof(f2240, plain, (~ spl66_21 | ~ spl66_23), inference(avatar_contradiction_clause, [], [f2239])).
fof(f2239, plain, ($false | (~ spl66_21 | ~ spl66_23)), inference(subsumption_resolution, [], [f2237, f284])).
fof(f2237, plain, ((e0 = e2) | (~ spl66_21 | ~ spl66_23)), inference(backward_demodulation, [], [f685, f677])).
fof(f2236, plain, (~ spl66_11 | ~ spl66_27), inference(avatar_split_clause, [], [f2235, f700, f632])).
fof(f2235, plain, (~ (e2 = op(e3, e1)) | ~ spl66_27), inference(backward_demodulation, [], [f246, f702])).
fof(f2220, plain, (~ spl66_40 | ~ spl66_44), inference(avatar_split_clause, [], [f2215, f772, f755])).
fof(f2215, plain, (~ (e3 = op(e1, e2)) | ~ spl66_44), inference(backward_demodulation, [], [f268, f774])).
fof(f2212, plain, (~ spl66_46 | ~ spl66_47), inference(avatar_contradiction_clause, [], [f2211])).
fof(f2211, plain, ($false | (~ spl66_46 | ~ spl66_47)), inference(subsumption_resolution, [], [f2210, f286])).
fof(f2210, plain, ((e1 = e2) | (~ spl66_46 | ~ spl66_47)), inference(backward_demodulation, [], [f787, f783])).
fof(f2201, plain, (~ spl66_53 | ~ spl66_57), inference(avatar_split_clause, [], [f2196, f828, f811])).
fof(f2196, plain, (~ (e0 = op(e0, e2)) | ~ spl66_57), inference(backward_demodulation, [], [f262, f830])).
fof(f2189, plain, (~ spl66_54 | ~ spl66_62), inference(avatar_split_clause, [], [f2183, f849, f815])).
fof(f2183, plain, (~ (e1 = op(e0, e2)) | ~ spl66_62), inference(backward_demodulation, [], [f260, f851])).
fof(f2161, plain, (~ spl66_34 | ~ spl66_18), inference(avatar_split_clause, [], [f2160, f662, f730])).
fof(f2160, plain, (~ (e1 = op(e1, e3)) | ~ spl66_18), inference(forward_demodulation, [], [f256, f664])).
fof(f2134, plain, (~ spl66_41 | ~ spl66_47 | spl66_137), inference(avatar_contradiction_clause, [], [f2133])).
fof(f2133, plain, ($false | (~ spl66_41 | ~ spl66_47 | spl66_137)), inference(subsumption_resolution, [], [f2132, f787])).
fof(f2132, plain, (~ (e2 = op(e1, e0)) | (~ spl66_41 | spl66_137)), inference(forward_demodulation, [], [f1462, f762])).
fof(f2113, plain, (~ spl66_31 | ~ spl66_23), inference(avatar_split_clause, [], [f2112, f683, f717])).
fof(f2112, plain, (~ (e2 = op(e2, e0)) | ~ spl66_23), inference(forward_demodulation, [], [f272, f685])).
fof(f2111, plain, (~ spl66_30 | ~ spl66_51 | ~ spl66_64 | spl66_136), inference(avatar_split_clause, [], [f1984, f1453, f857, f802, f713])).
fof(f1984, plain, (~ (e1 = op(e2, e0)) | (~ spl66_51 | ~ spl66_64 | spl66_136)), inference(backward_demodulation, [], [f1556, f804])).
fof(f1556, plain, (~ (e1 = op(op(e0, e3), e0)) | (~ spl66_64 | spl66_136)), inference(backward_demodulation, [], [f1455, f859])).
fof(f2110, plain, (~ spl66_31 | ~ spl66_47), inference(avatar_split_clause, [], [f2109, f785, f717])).
fof(f2109, plain, (~ (e2 = op(e2, e0)) | ~ spl66_47), inference(forward_demodulation, [], [f238, f787])).
fof(f2090, plain, (~ spl66_53 | ~ spl66_54), inference(avatar_contradiction_clause, [], [f2089])).
fof(f2089, plain, ($false | (~ spl66_53 | ~ spl66_54)), inference(subsumption_resolution, [], [f2088, f283])).
fof(f283, plain, ~ (e0 = e1), inference(cnf_transformation, [], [f4])).
fof(f2088, plain, ((e0 = e1) | (~ spl66_53 | ~ spl66_54)), inference(backward_demodulation, [], [f817, f813])).
fof(f2074, plain, (~ spl66_59 | ~ spl66_51), inference(avatar_split_clause, [], [f2073, f802, f836])).
fof(f2073, plain, (~ (e2 = op(e0, e1)) | ~ spl66_51), inference(forward_demodulation, [], [f263, f804])).
fof(f2070, plain, (~ spl66_55 | ~ spl66_51), inference(avatar_split_clause, [], [f2069, f802, f819])).
fof(f2069, plain, (~ (e2 = op(e0, e2)) | ~ spl66_51), inference(forward_demodulation, [], [f264, f804])).
fof(f2033, plain, (~ spl66_54 | spl66_146 | ~ spl66_147), inference(avatar_split_clause, [], [f2032, f1508, f1503, f815])).
fof(f2032, plain, (~ (e1 = op(e0, e2)) | (spl66_146 | ~ spl66_147)), inference(backward_demodulation, [], [f1505, f1509])).
fof(f2007, plain, (~ spl66_17 | ~ spl66_19), inference(avatar_contradiction_clause, [], [f2006])).
fof(f2006, plain, ($false | (~ spl66_17 | ~ spl66_19)), inference(subsumption_resolution, [], [f2005, f284])).
fof(f2005, plain, ((e0 = e2) | (~ spl66_17 | ~ spl66_19)), inference(backward_demodulation, [], [f668, f660])).
fof(f1999, plain, (~ spl66_9 | ~ spl66_40 | ~ spl66_43 | spl66_142), inference(avatar_split_clause, [], [f1996, f1483, f768, f755, f624])).
fof(f1996, plain, (~ (e0 = op(e3, e1)) | (~ spl66_40 | ~ spl66_43 | spl66_142)), inference(backward_demodulation, [], [f1804, f757])).
fof(f1804, plain, (~ (e0 = op(op(e1, e2), e1)) | (~ spl66_43 | spl66_142)), inference(backward_demodulation, [], [f1485, f770])).
fof(f1995, plain, (~ spl66_41 | ~ spl66_42), inference(avatar_contradiction_clause, [], [f1994])).
fof(f1994, plain, ($false | (~ spl66_41 | ~ spl66_42)), inference(subsumption_resolution, [], [f1993, f283])).
fof(f1993, plain, ((e0 = e1) | (~ spl66_41 | ~ spl66_42)), inference(backward_demodulation, [], [f766, f762])).
fof(f1992, plain, (~ spl66_42 | ~ spl66_43), inference(avatar_contradiction_clause, [], [f1991])).
fof(f1991, plain, ($false | (~ spl66_42 | ~ spl66_43)), inference(subsumption_resolution, [], [f1990, f286])).
fof(f1990, plain, ((e1 = e2) | (~ spl66_42 | ~ spl66_43)), inference(backward_demodulation, [], [f770, f766])).
fof(f1983, plain, (~ spl66_61 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f1982])).
fof(f1982, plain, ($false | (~ spl66_61 | ~ spl66_64)), inference(subsumption_resolution, [], [f1981, f285])).
fof(f1981, plain, ((e0 = e3) | (~ spl66_61 | ~ spl66_64)), inference(backward_demodulation, [], [f859, f847])).
fof(f1957, plain, (~ spl66_50 | ~ spl66_34), inference(avatar_split_clause, [], [f1956, f730, f798])).
fof(f1956, plain, (~ (e1 = op(e0, e3)) | ~ spl66_34), inference(forward_demodulation, [], [f253, f732])).
fof(f1955, plain, (~ spl66_46 | ~ spl66_34), inference(avatar_split_clause, [], [f1954, f730, f781])).
fof(f1954, plain, (~ (e1 = op(e1, e0)) | ~ spl66_34), inference(forward_demodulation, [], [f267, f732])).
fof(f1953, plain, (~ spl66_42 | ~ spl66_34), inference(avatar_split_clause, [], [f1952, f730, f764])).
fof(f1952, plain, (~ (e1 = op(e1, e1)) | ~ spl66_34), inference(forward_demodulation, [], [f269, f732])).
fof(f1921, plain, (~ spl66_33 | ~ spl66_34), inference(avatar_contradiction_clause, [], [f1920])).
fof(f1920, plain, ($false | (~ spl66_33 | ~ spl66_34)), inference(subsumption_resolution, [], [f1919, f283])).
fof(f1919, plain, ((e0 = e1) | (~ spl66_33 | ~ spl66_34)), inference(forward_demodulation, [], [f732, f728])).
fof(f728, plain, ((e0 = op(e1, e3)) | ~ spl66_33), inference(avatar_component_clause, [], [f726])).
fof(f1884, plain, (~ spl66_21 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f1883])).
fof(f1883, plain, ($false | (~ spl66_21 | ~ spl66_24)), inference(subsumption_resolution, [], [f1882, f285])).
fof(f1882, plain, ((e0 = e3) | (~ spl66_21 | ~ spl66_24)), inference(forward_demodulation, [], [f689, f677])).
fof(f1858, plain, (~ spl66_1 | ~ spl66_9), inference(avatar_split_clause, [], [f1857, f624, f590])).
fof(f1857, plain, (~ (e0 = op(e3, e3)) | ~ spl66_9), inference(forward_demodulation, [], [f281, f626])).
fof(f1853, plain, (~ spl66_1 | ~ spl66_49), inference(avatar_split_clause, [], [f1852, f794, f590])).
fof(f1852, plain, (~ (e0 = op(e3, e3)) | ~ spl66_49), inference(forward_demodulation, [], [f255, f796])).
fof(f1837, plain, (~ spl66_18 | ~ spl66_19), inference(avatar_contradiction_clause, [], [f1836])).
fof(f1836, plain, ($false | (~ spl66_18 | ~ spl66_19)), inference(subsumption_resolution, [], [f1835, f286])).
fof(f1835, plain, ((e1 = e2) | (~ spl66_18 | ~ spl66_19)), inference(backward_demodulation, [], [f668, f664])).
fof(f1833, plain, (~ spl66_21 | ~ spl66_22), inference(avatar_contradiction_clause, [], [f1832])).
fof(f1832, plain, ($false | (~ spl66_21 | ~ spl66_22)), inference(subsumption_resolution, [], [f1831, f283])).
fof(f1831, plain, ((e0 = e1) | (~ spl66_21 | ~ spl66_22)), inference(backward_demodulation, [], [f681, f677])).
fof(f1777, plain, (~ spl66_62 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f1776])).
fof(f1776, plain, ($false | (~ spl66_62 | ~ spl66_64)), inference(subsumption_resolution, [], [f1775, f287])).
fof(f1775, plain, ((e1 = e3) | (~ spl66_62 | ~ spl66_64)), inference(backward_demodulation, [], [f859, f851])).
fof(f1767, plain, (~ spl66_28 | ~ spl66_22 | spl66_145), inference(avatar_split_clause, [], [f1766, f1499, f679, f704])).
fof(f1766, plain, (~ (e3 = op(e2, e1)) | (~ spl66_22 | spl66_145)), inference(forward_demodulation, [], [f1501, f681])).
fof(f1758, plain, (~ spl66_53 | ~ spl66_49), inference(avatar_split_clause, [], [f1757, f794, f811])).
fof(f1757, plain, (~ (e0 = op(e0, e2)) | ~ spl66_49), inference(forward_demodulation, [], [f264, f796])).
fof(f1743, plain, (~ spl66_43 | ~ spl66_35), inference(avatar_split_clause, [], [f1742, f734, f768])).
fof(f1742, plain, (~ (e2 = op(e1, e1)) | ~ spl66_35), inference(forward_demodulation, [], [f269, f736])).
fof(f1741, plain, (~ spl66_38 | ~ spl66_22), inference(avatar_split_clause, [], [f1740, f679, f747])).
fof(f1740, plain, (~ (e1 = op(e1, e2)) | ~ spl66_22), inference(forward_demodulation, [], [f250, f681])).
fof(f1726, plain, (~ spl66_26 | ~ spl66_22), inference(avatar_split_clause, [], [f1725, f679, f696])).
fof(f1725, plain, (~ (e1 = op(e2, e1)) | ~ spl66_22), inference(forward_demodulation, [], [f274, f681])).
fof(f1724, plain, (~ spl66_27 | ~ spl66_19), inference(avatar_split_clause, [], [f1723, f666, f700])).
fof(f1723, plain, (~ (e2 = op(e2, e1)) | ~ spl66_19), inference(forward_demodulation, [], [f275, f668])).
fof(f1716, plain, (~ spl66_13 | ~ spl66_1), inference(avatar_split_clause, [], [f1715, f590, f641])).
fof(f1715, plain, (~ (e0 = op(e3, e0)) | ~ spl66_1), inference(forward_demodulation, [], [f279, f592])).
fof(f592, plain, ((e0 = op(e3, e3)) | ~ spl66_1), inference(avatar_component_clause, [], [f590])).
fof(f1708, plain, (~ spl66_1 | ~ spl66_4), inference(avatar_contradiction_clause, [], [f1707])).
fof(f1707, plain, ($false | (~ spl66_1 | ~ spl66_4)), inference(subsumption_resolution, [], [f1706, f285])).
fof(f1706, plain, ((e0 = e3) | (~ spl66_1 | ~ spl66_4)), inference(forward_demodulation, [], [f604, f592])).
fof(f1702, plain, (~ spl66_2 | ~ spl66_3), inference(avatar_contradiction_clause, [], [f1701])).
fof(f1701, plain, ($false | (~ spl66_2 | ~ spl66_3)), inference(subsumption_resolution, [], [f1700, f286])).
fof(f1700, plain, ((e1 = e2) | (~ spl66_2 | ~ spl66_3)), inference(backward_demodulation, [], [f600, f596])).
fof(f1699, plain, (~ spl66_3 | ~ spl66_4), inference(avatar_contradiction_clause, [], [f1698])).
fof(f1698, plain, ($false | (~ spl66_3 | ~ spl66_4)), inference(subsumption_resolution, [], [f1697, f288])).
fof(f1697, plain, ((e2 = e3) | (~ spl66_3 | ~ spl66_4)), inference(backward_demodulation, [], [f604, f600])).
fof(f1664, plain, (~ spl66_22 | ~ spl66_23), inference(avatar_contradiction_clause, [], [f1663])).
fof(f1663, plain, ($false | (~ spl66_22 | ~ spl66_23)), inference(subsumption_resolution, [], [f1662, f286])).
fof(f1662, plain, ((e1 = e2) | (~ spl66_22 | ~ spl66_23)), inference(backward_demodulation, [], [f685, f681])).
fof(f1661, plain, (~ spl66_23 | ~ spl66_24), inference(avatar_contradiction_clause, [], [f1660])).
fof(f1660, plain, ($false | (~ spl66_23 | ~ spl66_24)), inference(subsumption_resolution, [], [f1659, f288])).
fof(f1659, plain, ((e2 = e3) | (~ spl66_23 | ~ spl66_24)), inference(backward_demodulation, [], [f689, f685])).
fof(f1658, plain, (~ spl66_17 | ~ spl66_24 | spl66_147), inference(avatar_split_clause, [], [f1654, f1508, f687, f658])).
fof(f1654, plain, (~ (e0 = op(e2, e3)) | (~ spl66_24 | spl66_147)), inference(backward_demodulation, [], [f1510, f689])).
fof(f1657, plain, (~ spl66_20 | ~ spl66_24), inference(avatar_split_clause, [], [f1653, f687, f670])).
fof(f1653, plain, (~ (e3 = op(e2, e3)) | ~ spl66_24), inference(backward_demodulation, [], [f276, f689])).
fof(f1656, plain, (~ spl66_8 | ~ spl66_24), inference(avatar_split_clause, [], [f1652, f687, f619])).
fof(f1652, plain, (~ (e3 = op(e3, e2)) | ~ spl66_24), inference(backward_demodulation, [], [f252, f689])).
fof(f1650, plain, (~ spl66_21 | ~ spl66_25), inference(avatar_split_clause, [], [f1647, f692, f675])).
fof(f1647, plain, (~ (e0 = op(e2, e2)) | ~ spl66_25), inference(backward_demodulation, [], [f274, f694])).
fof(f1649, plain, (~ spl66_9 | ~ spl66_25), inference(avatar_split_clause, [], [f1646, f692, f624])).
fof(f1646, plain, (~ (e0 = op(e3, e1)) | ~ spl66_25), inference(backward_demodulation, [], [f246, f694])).
fof(f1619, plain, (~ spl66_33 | ~ spl66_41), inference(avatar_split_clause, [], [f1610, f760, f726])).
fof(f1610, plain, (~ (e0 = op(e1, e3)) | ~ spl66_41), inference(backward_demodulation, [], [f269, f762])).
fof(f1618, plain, (~ spl66_37 | ~ spl66_41), inference(avatar_split_clause, [], [f1609, f760, f743])).
fof(f1609, plain, (~ (e0 = op(e1, e2)) | ~ spl66_41), inference(backward_demodulation, [], [f268, f762])).
fof(f1604, plain, (~ spl66_41 | ~ spl66_45), inference(avatar_split_clause, [], [f1599, f777, f760])).
fof(f1599, plain, (~ (e0 = op(e1, e1)) | ~ spl66_45), inference(backward_demodulation, [], [f265, f779])).
fof(f1595, plain, (~ spl66_49 | ~ spl66_50), inference(avatar_contradiction_clause, [], [f1594])).
fof(f1594, plain, ($false | (~ spl66_49 | ~ spl66_50)), inference(subsumption_resolution, [], [f1593, f283])).
fof(f1593, plain, ((e0 = e1) | (~ spl66_49 | ~ spl66_50)), inference(backward_demodulation, [], [f800, f796])).
fof(f1583, plain, (~ spl66_5 | ~ spl66_53), inference(avatar_split_clause, [], [f1579, f811, f607])).
fof(f1579, plain, (~ (e0 = op(e3, e2)) | ~ spl66_53), inference(backward_demodulation, [], [f249, f813])).
fof(f1573, plain, (~ spl66_25 | ~ spl66_57), inference(avatar_split_clause, [], [f1568, f828, f692])).
fof(f1568, plain, (~ (e0 = op(e2, e1)) | ~ spl66_57), inference(backward_demodulation, [], [f242, f830])).
fof(f1566, plain, (~ spl66_63 | ~ spl66_64), inference(avatar_contradiction_clause, [], [f1565])).
fof(f1565, plain, ($false | (~ spl66_63 | ~ spl66_64)), inference(subsumption_resolution, [], [f1564, f288])).
fof(f1564, plain, ((e2 = e3) | (~ spl66_63 | ~ spl66_64)), inference(backward_demodulation, [], [f859, f855])).
fof(f1563, plain, (~ spl66_50 | ~ spl66_64 | spl66_135), inference(avatar_split_clause, [], [f1555, f1448, f857, f798])).
fof(f1555, plain, (~ (e1 = op(e0, e3)) | (~ spl66_64 | spl66_135)), inference(backward_demodulation, [], [f1450, f859])).
fof(f1560, plain, (~ spl66_60 | ~ spl66_64), inference(avatar_split_clause, [], [f1552, f857, f840])).
fof(f1552, plain, (~ (e3 = op(e0, e1)) | ~ spl66_64), inference(backward_demodulation, [], [f259, f859])).
fof(f1558, plain, (~ spl66_32 | ~ spl66_64), inference(avatar_split_clause, [], [f1550, f857, f721])).
fof(f1550, plain, (~ (e3 = op(e2, e0)) | ~ spl66_64), inference(backward_demodulation, [], [f236, f859])).
fof(f1547, plain, (~ spl66_3 | ~ spl66_153 | ~ spl66_152), inference(avatar_split_clause, [], [f587, f1533, f1538, f598])).
fof(f587, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, e3))), inference(cnf_transformation, [], [f52])).
fof(f52, plain, (~ (e1 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e2 = op(e3, e3))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ~ ((e1 = op(op(e3, op(e3, e3)), e3)) & (e0 = op(e3, op(e3, e3))) & (e2 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax28)).
fof(f1541, plain, (~ spl66_2 | ~ spl66_153 | ~ spl66_150), inference(avatar_split_clause, [], [f585, f1524, f1538, f594])).
fof(f585, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(e3, e3))), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (~ (e2 = op(op(e3, op(e3, e3)), e3)) | ~ (e0 = op(e3, op(e3, e3))) | ~ (e1 = op(e3, e3))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ ((e2 = op(op(e3, op(e3, e3)), e3)) & (e0 = op(e3, op(e3, e3))) & (e1 = op(e3, e3))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax26)).
fof(f1518, plain, (~ spl66_24 | ~ spl66_143 | ~ spl66_148), inference(avatar_split_clause, [], [f582, f1513, f1490, f687])).
fof(f582, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(cnf_transformation, [], [f47])).
fof(f47, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ~ ((e0 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, op(e2, e2))) & (e3 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax23)).
fof(f1517, plain, (~ spl66_24 | ~ spl66_147 | ~ spl66_146), inference(avatar_split_clause, [], [f581, f1503, f1508, f687])).
fof(f581, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e3 = op(e2, e2))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ~ ((e1 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, op(e2, e2))) & (e3 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax22)).
fof(f1516, plain, (~ spl66_22 | ~ spl66_145 | ~ spl66_148), inference(avatar_split_clause, [], [f580, f1513, f1499, f679])).
fof(f580, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(cnf_transformation, [], [f45])).
fof(f45, plain, (~ (e0 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ~ ((e0 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, op(e2, e2))) & (e1 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax21)).
fof(f1511, plain, (~ spl66_22 | ~ spl66_147 | ~ spl66_144), inference(avatar_split_clause, [], [f579, f1494, f1508, f679])).
fof(f579, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(cnf_transformation, [], [f44])).
fof(f44, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e0 = op(e2, op(e2, e2))) | ~ (e1 = op(e2, e2))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ~ ((e3 = op(op(e2, op(e2, e2)), e2)) & (e0 = op(e2, op(e2, e2))) & (e1 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax20)).
fof(f1506, plain, (~ spl66_21 | ~ spl66_145 | ~ spl66_146), inference(avatar_split_clause, [], [f578, f1503, f1499, f675])).
fof(f578, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(e2, e2))), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (~ (e1 = op(op(e2, op(e2, e2)), e2)) | ~ (e3 = op(e2, op(e2, e2))) | ~ (e0 = op(e2, e2))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ~ ((e1 = op(op(e2, op(e2, e2)), e2)) & (e3 = op(e2, op(e2, e2))) & (e0 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax19)).
fof(f1497, plain, (~ spl66_21 | ~ spl66_143 | ~ spl66_144), inference(avatar_split_clause, [], [f577, f1494, f1490, f675])).
fof(f577, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(e2, e2))), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (~ (e3 = op(op(e2, op(e2, e2)), e2)) | ~ (e1 = op(e2, op(e2, e2))) | ~ (e0 = op(e2, e2))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ~ ((e3 = op(op(e2, op(e2, e2)), e2)) & (e1 = op(e2, op(e2, e2))) & (e0 = op(e2, e2))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax18)).
fof(f1488, plain, (~ spl66_44 | ~ spl66_137 | ~ spl66_142), inference(avatar_split_clause, [], [f576, f1483, f1460, f772])).
fof(f576, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(e1, e1))), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e3 = op(e1, e1))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ~ ((e0 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, op(e1, e1))) & (e3 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax17)).
fof(f1486, plain, (~ spl66_43 | ~ spl66_139 | ~ spl66_142), inference(avatar_split_clause, [], [f574, f1483, f1469, f768])).
fof(f574, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(e1, e1))), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (~ (e0 = op(op(e1, op(e1, e1)), e1)) | ~ (e3 = op(e1, op(e1, e1))) | ~ (e2 = op(e1, e1))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ((e0 = op(op(e1, op(e1, e1)), e1)) & (e3 = op(e1, op(e1, e1))) & (e2 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax15)).
fof(f1467, plain, (~ spl66_41 | ~ spl66_137 | ~ spl66_138), inference(avatar_split_clause, [], [f571, f1464, f1460, f760])).
fof(f571, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1))), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (~ (e3 = op(op(e1, op(e1, e1)), e1)) | ~ (e2 = op(e1, op(e1, e1))) | ~ (e0 = op(e1, e1))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ~ ((e3 = op(op(e1, op(e1, e1)), e1)) & (e2 = op(e1, op(e1, e1))) & (e0 = op(e1, e1))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax12)).
fof(f1458, plain, (~ spl66_64 | ~ spl66_131 | ~ spl66_136), inference(avatar_split_clause, [], [f570, f1453, f1430, f857])).
fof(f570, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (e2 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e3)), inference(cnf_transformation, [], [f35])).
fof(f35, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (e2 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e3)), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ~ ((e1 = op(op(e0, op(e0, e0)), e0)) & (e2 = op(e0, op(e0, e0))) & (op(e0, e0) = e3)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax11)).
fof(f1457, plain, (~ spl66_64 | ~ spl66_135 | ~ spl66_134), inference(avatar_split_clause, [], [f569, f1443, f1448, f857])).
fof(f569, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (e1 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e3)), inference(cnf_transformation, [], [f34])).
fof(f34, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (e1 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e3)), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ((e2 = op(op(e0, op(e0, e0)), e0)) & (e1 = op(e0, op(e0, e0))) & (op(e0, e0) = e3)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax10)).
fof(f1456, plain, (~ spl66_63 | ~ spl66_133 | ~ spl66_136), inference(avatar_split_clause, [], [f568, f1453, f1439, f853])).
fof(f568, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e2)), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (~ (e1 = op(op(e0, op(e0, e0)), e0)) | ~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e2)), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ~ ((e1 = op(op(e0, op(e0, e0)), e0)) & (e3 = op(e0, op(e0, e0))) & (op(e0, e0) = e2)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax9)).
fof(f1451, plain, (~ spl66_63 | ~ spl66_135 | ~ spl66_132), inference(avatar_split_clause, [], [f567, f1434, f1448, f853])).
fof(f567, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (e1 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e2)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (e1 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e2)), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ~ ((e3 = op(op(e0, op(e0, e0)), e0)) & (e1 = op(e0, op(e0, e0))) & (op(e0, e0) = e2)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax8)).
fof(f1446, plain, (~ spl66_62 | ~ spl66_133 | ~ spl66_134), inference(avatar_split_clause, [], [f566, f1443, f1439, f849])).
fof(f566, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1)), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (~ (e2 = op(op(e0, op(e0, e0)), e0)) | ~ (e3 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1)), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((e2 = op(op(e0, op(e0, e0)), e0)) & (e3 = op(e0, op(e0, e0))) & (op(e0, e0) = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax7)).
fof(f1437, plain, (~ spl66_62 | ~ spl66_131 | ~ spl66_132), inference(avatar_split_clause, [], [f565, f1434, f1430, f849])).
fof(f565, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (e2 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ (e3 = op(op(e0, op(e0, e0)), e0)) | ~ (e2 = op(e0, op(e0, e0))) | ~ (op(e0, e0) = e1)), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((e3 = op(op(e0, op(e0, e0)), e0)) & (e2 = op(e0, op(e0, e0))) & (op(e0, e0) = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax6)).
fof(f1428, plain, (spl66_130 | spl66_129 | spl66_128 | spl66_4), inference(avatar_split_clause, [], [f556, f602, f1398, f1406, f1414])).
fof(f1414, plain, (spl66_130 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl66_130])])).
fof(f1406, plain, (spl66_129 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl66_129])])).
fof(f1398, plain, (spl66_128 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl66_128])])).
fof(f556, plain, ((e3 = op(e3, e3)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f120, plain, (((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3) & ((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2))) & ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1))) & ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3)) & (e3 = op(e3, e3))) | sP2 | sP1 | sP0)), inference(definition_folding, [], [f5, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54])).
fof(f54, plain, ((((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e54])).
fof(e54, plain, (sP0 <=> (((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f55, plain, ((((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e1 = op(e1, e1))) | ~ sP1), inference(usedef, [], [e55])).
fof(e55, plain, (sP1 <=> (((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f56, plain, ((((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e2 = op(e2, e2))) | ~ sP2), inference(usedef, [], [e56])).
fof(e56, plain, (sP2 <=> (((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f57, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(usedef, [], [e57])).
fof(e57, plain, (sP3 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f58, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP4), inference(usedef, [], [e58])).
fof(e58, plain, (sP4 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f59, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP5), inference(usedef, [], [e59])).
fof(e59, plain, (sP5 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f60, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP6), inference(usedef, [], [e60])).
fof(e60, plain, (sP6 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f61, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP7), inference(usedef, [], [e61])).
fof(e61, plain, (sP7 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f62, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP8), inference(usedef, [], [e62])).
fof(e62, plain, (sP8 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f63, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP9), inference(usedef, [], [e63])).
fof(e63, plain, (sP9 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f64, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP10), inference(usedef, [], [e64])).
fof(e64, plain, (sP10 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f65, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP11), inference(usedef, [], [e65])).
fof(e65, plain, (sP11 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f66, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP12), inference(usedef, [], [e66])).
fof(e66, plain, (sP12 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f67, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP13), inference(usedef, [], [e67])).
fof(e67, plain, (sP13 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f68, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP14), inference(usedef, [], [e68])).
fof(e68, plain, (sP14 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f69, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP15), inference(usedef, [], [e69])).
fof(e69, plain, (sP15 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f70, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP16), inference(usedef, [], [e70])).
fof(e70, plain, (sP16 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f71, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP17), inference(usedef, [], [e71])).
fof(e71, plain, (sP17 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f72, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP18), inference(usedef, [], [e72])).
fof(e72, plain, (sP18 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f73, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(usedef, [], [e73])).
fof(e73, plain, (sP19 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f74, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP20), inference(usedef, [], [e74])).
fof(e74, plain, (sP20 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f75, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP21), inference(usedef, [], [e75])).
fof(e75, plain, (sP21 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f76, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP22), inference(usedef, [], [e76])).
fof(e76, plain, (sP22 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f77, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(usedef, [], [e77])).
fof(e77, plain, (sP23 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f78, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP24), inference(usedef, [], [e78])).
fof(e78, plain, (sP24 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f79, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP25), inference(usedef, [], [e79])).
fof(e79, plain, (sP25 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f80, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP26), inference(usedef, [], [e80])).
fof(e80, plain, (sP26 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f81, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(usedef, [], [e81])).
fof(e81, plain, (sP27 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f82, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP28), inference(usedef, [], [e82])).
fof(e82, plain, (sP28 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f83, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP29), inference(usedef, [], [e83])).
fof(e83, plain, (sP29 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f84, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP30), inference(usedef, [], [e84])).
fof(e84, plain, (sP30 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f85, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(usedef, [], [e85])).
fof(e85, plain, (sP31 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f86, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP32), inference(usedef, [], [e86])).
fof(e86, plain, (sP32 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f87, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP33), inference(usedef, [], [e87])).
fof(e87, plain, (sP33 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f88, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP34), inference(usedef, [], [e88])).
fof(e88, plain, (sP34 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f89, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(usedef, [], [e89])).
fof(e89, plain, (sP35 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f90, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP36), inference(usedef, [], [e90])).
fof(e90, plain, (sP36 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f91, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP37), inference(usedef, [], [e91])).
fof(e91, plain, (sP37 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f92, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP38), inference(usedef, [], [e92])).
fof(e92, plain, (sP38 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f93, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(usedef, [], [e93])).
fof(e93, plain, (sP39 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f94, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP40), inference(usedef, [], [e94])).
fof(e94, plain, (sP40 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f95, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP41), inference(usedef, [], [e95])).
fof(e95, plain, (sP41 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f96, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP42), inference(usedef, [], [e96])).
fof(e96, plain, (sP42 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f97, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(usedef, [], [e97])).
fof(e97, plain, (sP43 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f98, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP44), inference(usedef, [], [e98])).
fof(e98, plain, (sP44 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f99, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP45), inference(usedef, [], [e99])).
fof(e99, plain, (sP45 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f100, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP46), inference(usedef, [], [e100])).
fof(e100, plain, (sP46 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f101, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(usedef, [], [e101])).
fof(e101, plain, (sP47 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f102, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP48), inference(usedef, [], [e102])).
fof(e102, plain, (sP48 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f103, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP49), inference(usedef, [], [e103])).
fof(e103, plain, (sP49 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f104, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP50), inference(usedef, [], [e104])).
fof(e104, plain, (sP50 <=> (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f105, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(usedef, [], [e105])).
fof(e105, plain, (sP51 <=> (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f106, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP52), inference(usedef, [], [e106])).
fof(e106, plain, (sP52 <=> (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f107, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP53), inference(usedef, [], [e107])).
fof(e107, plain, (sP53 <=> (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f108, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP54), inference(usedef, [], [e108])).
fof(e108, plain, (sP54 <=> (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f109, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(usedef, [], [e109])).
fof(e109, plain, (sP55 <=> (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f110, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP56), inference(usedef, [], [e110])).
fof(e110, plain, (sP56 <=> (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f111, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP57), inference(usedef, [], [e111])).
fof(e111, plain, (sP57 <=> (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f112, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP58), inference(usedef, [], [e112])).
fof(e112, plain, (sP58 <=> (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f113, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(usedef, [], [e113])).
fof(e113, plain, (sP59 <=> (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f114, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP60), inference(usedef, [], [e114])).
fof(e114, plain, (sP60 <=> (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f115, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP61), inference(usedef, [], [e115])).
fof(e115, plain, (sP61 <=> (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f116, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP62), inference(usedef, [], [e116])).
fof(e116, plain, (sP62 <=> (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f117, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP63), inference(usedef, [], [e117])).
fof(e117, plain, (sP63 <=> (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f118, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP64), inference(usedef, [], [e118])).
fof(e118, plain, (sP64 <=> (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f119, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP65), inference(usedef, [], [e119])).
fof(e119, plain, (sP65 <=> (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f5, plain, (((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | (~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | (~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | (~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | (~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | (~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)))) & ((((e3 = op(e3, e3)) | ~ (e3 = op(e3, e3))) & ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2))) & ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1))) & ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3)) & (e3 = op(e3, e3))) | (((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e2 = op(e2, e2))) | (((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e1 = op(e1, e1))) | (((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0))))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax5)).
fof(f1427, plain, (spl66_130 | spl66_129 | spl66_128 | ~ spl66_64 | spl66_49), inference(avatar_split_clause, [], [f557, f794, f857, f1398, f1406, f1414])).
fof(f557, plain, ((e0 = op(e0, e3)) | ~ (op(e0, e0) = e3) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1426, plain, (spl66_130 | spl66_129 | spl66_128 | ~ spl66_44 | spl66_34), inference(avatar_split_clause, [], [f558, f730, f772, f1398, f1406, f1414])).
fof(f558, plain, ((e1 = op(e1, e3)) | ~ (e3 = op(e1, e1)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1425, plain, (spl66_130 | spl66_129 | spl66_128 | ~ spl66_24 | spl66_19), inference(avatar_split_clause, [], [f559, f666, f687, f1398, f1406, f1414])).
fof(f559, plain, ((e2 = op(e2, e3)) | ~ (e3 = op(e2, e2)) | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f120])).
fof(f1424, plain, (spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | spl66_79 | spl66_78 | spl66_77 | spl66_76 | spl66_75 | spl66_74 | spl66_73 | spl66_72 | spl66_71 | spl66_70 | spl66_69 | spl66_68 | spl66_67 | spl66_66 | spl66_65 | spl66_4), inference(avatar_split_clause, [], [f561, f602, f894, f902, f910, f918, f926, f934, f942, f950, f958, f966, f974, f982, f990, f998, f1006, f1014, f1022, f1030, f1038, f1046, f1054, f1062, f1070, f1078, f1086, f1094, f1102, f1110, f1118, f1126, f1134, f1142, f1150, f1158, f1166, f1174, f1182, f1190, f1198, f1206, f1214, f1222, f1230, f1238, f1246, f1254, f1262, f1270, f1278, f1286, f1294, f1302, f1310, f1318, f1326, f1334, f1342, f1350, f1358, f1366, f1374, f1382, f1390])).
fof(f1390, plain, (spl66_127 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl66_127])])).
fof(f1382, plain, (spl66_126 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl66_126])])).
fof(f1374, plain, (spl66_125 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl66_125])])).
fof(f1366, plain, (spl66_124 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl66_124])])).
fof(f1358, plain, (spl66_123 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl66_123])])).
fof(f1350, plain, (spl66_122 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl66_122])])).
fof(f1342, plain, (spl66_121 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl66_121])])).
fof(f1334, plain, (spl66_120 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl66_120])])).
fof(f1326, plain, (spl66_119 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl66_119])])).
fof(f1318, plain, (spl66_118 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl66_118])])).
fof(f1310, plain, (spl66_117 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl66_117])])).
fof(f1302, plain, (spl66_116 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl66_116])])).
fof(f1294, plain, (spl66_115 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl66_115])])).
fof(f1286, plain, (spl66_114 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl66_114])])).
fof(f1278, plain, (spl66_113 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl66_113])])).
fof(f1270, plain, (spl66_112 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl66_112])])).
fof(f1262, plain, (spl66_111 <=> sP19), introduced(avatar_definition, [new_symbols(naming, [spl66_111])])).
fof(f1254, plain, (spl66_110 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl66_110])])).
fof(f1246, plain, (spl66_109 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl66_109])])).
fof(f1238, plain, (spl66_108 <=> sP22), introduced(avatar_definition, [new_symbols(naming, [spl66_108])])).
fof(f1230, plain, (spl66_107 <=> sP23), introduced(avatar_definition, [new_symbols(naming, [spl66_107])])).
fof(f1222, plain, (spl66_106 <=> sP24), introduced(avatar_definition, [new_symbols(naming, [spl66_106])])).
fof(f1214, plain, (spl66_105 <=> sP25), introduced(avatar_definition, [new_symbols(naming, [spl66_105])])).
fof(f1206, plain, (spl66_104 <=> sP26), introduced(avatar_definition, [new_symbols(naming, [spl66_104])])).
fof(f1198, plain, (spl66_103 <=> sP27), introduced(avatar_definition, [new_symbols(naming, [spl66_103])])).
fof(f1190, plain, (spl66_102 <=> sP28), introduced(avatar_definition, [new_symbols(naming, [spl66_102])])).
fof(f1182, plain, (spl66_101 <=> sP29), introduced(avatar_definition, [new_symbols(naming, [spl66_101])])).
fof(f1174, plain, (spl66_100 <=> sP30), introduced(avatar_definition, [new_symbols(naming, [spl66_100])])).
fof(f1166, plain, (spl66_99 <=> sP31), introduced(avatar_definition, [new_symbols(naming, [spl66_99])])).
fof(f1158, plain, (spl66_98 <=> sP32), introduced(avatar_definition, [new_symbols(naming, [spl66_98])])).
fof(f1150, plain, (spl66_97 <=> sP33), introduced(avatar_definition, [new_symbols(naming, [spl66_97])])).
fof(f1142, plain, (spl66_96 <=> sP34), introduced(avatar_definition, [new_symbols(naming, [spl66_96])])).
fof(f1134, plain, (spl66_95 <=> sP35), introduced(avatar_definition, [new_symbols(naming, [spl66_95])])).
fof(f1126, plain, (spl66_94 <=> sP36), introduced(avatar_definition, [new_symbols(naming, [spl66_94])])).
fof(f1118, plain, (spl66_93 <=> sP37), introduced(avatar_definition, [new_symbols(naming, [spl66_93])])).
fof(f1110, plain, (spl66_92 <=> sP38), introduced(avatar_definition, [new_symbols(naming, [spl66_92])])).
fof(f1102, plain, (spl66_91 <=> sP39), introduced(avatar_definition, [new_symbols(naming, [spl66_91])])).
fof(f1094, plain, (spl66_90 <=> sP40), introduced(avatar_definition, [new_symbols(naming, [spl66_90])])).
fof(f1086, plain, (spl66_89 <=> sP41), introduced(avatar_definition, [new_symbols(naming, [spl66_89])])).
fof(f1078, plain, (spl66_88 <=> sP42), introduced(avatar_definition, [new_symbols(naming, [spl66_88])])).
fof(f1070, plain, (spl66_87 <=> sP43), introduced(avatar_definition, [new_symbols(naming, [spl66_87])])).
fof(f1062, plain, (spl66_86 <=> sP44), introduced(avatar_definition, [new_symbols(naming, [spl66_86])])).
fof(f1054, plain, (spl66_85 <=> sP45), introduced(avatar_definition, [new_symbols(naming, [spl66_85])])).
fof(f1046, plain, (spl66_84 <=> sP46), introduced(avatar_definition, [new_symbols(naming, [spl66_84])])).
fof(f1038, plain, (spl66_83 <=> sP47), introduced(avatar_definition, [new_symbols(naming, [spl66_83])])).
fof(f1030, plain, (spl66_82 <=> sP48), introduced(avatar_definition, [new_symbols(naming, [spl66_82])])).
fof(f1022, plain, (spl66_81 <=> sP49), introduced(avatar_definition, [new_symbols(naming, [spl66_81])])).
fof(f1014, plain, (spl66_80 <=> sP50), introduced(avatar_definition, [new_symbols(naming, [spl66_80])])).
fof(f1006, plain, (spl66_79 <=> sP51), introduced(avatar_definition, [new_symbols(naming, [spl66_79])])).
fof(f998, plain, (spl66_78 <=> sP52), introduced(avatar_definition, [new_symbols(naming, [spl66_78])])).
fof(f990, plain, (spl66_77 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl66_77])])).
fof(f982, plain, (spl66_76 <=> sP54), introduced(avatar_definition, [new_symbols(naming, [spl66_76])])).
fof(f974, plain, (spl66_75 <=> sP55), introduced(avatar_definition, [new_symbols(naming, [spl66_75])])).
fof(f966, plain, (spl66_74 <=> sP56), introduced(avatar_definition, [new_symbols(naming, [spl66_74])])).
fof(f958, plain, (spl66_73 <=> sP57), introduced(avatar_definition, [new_symbols(naming, [spl66_73])])).
fof(f950, plain, (spl66_72 <=> sP58), introduced(avatar_definition, [new_symbols(naming, [spl66_72])])).
fof(f942, plain, (spl66_71 <=> sP59), introduced(avatar_definition, [new_symbols(naming, [spl66_71])])).
fof(f934, plain, (spl66_70 <=> sP60), introduced(avatar_definition, [new_symbols(naming, [spl66_70])])).
fof(f926, plain, (spl66_69 <=> sP61), introduced(avatar_definition, [new_symbols(naming, [spl66_69])])).
fof(f918, plain, (spl66_68 <=> sP62), introduced(avatar_definition, [new_symbols(naming, [spl66_68])])).
fof(f910, plain, (spl66_67 <=> sP63), introduced(avatar_definition, [new_symbols(naming, [spl66_67])])).
fof(f902, plain, (spl66_66 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl66_66])])).
fof(f894, plain, (spl66_65 <=> sP65), introduced(avatar_definition, [new_symbols(naming, [spl66_65])])).
fof(f561, plain, ((e3 = op(e3, e3)) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f120])).
fof(f1423, plain, (spl66_127 | spl66_126 | spl66_125 | spl66_124 | spl66_123 | spl66_122 | spl66_121 | spl66_120 | spl66_119 | spl66_118 | spl66_117 | spl66_116 | spl66_115 | spl66_114 | spl66_113 | spl66_112 | spl66_111 | spl66_110 | spl66_109 | spl66_108 | spl66_107 | spl66_106 | spl66_105 | spl66_104 | spl66_103 | spl66_102 | spl66_101 | spl66_100 | spl66_99 | spl66_98 | spl66_97 | spl66_96 | spl66_95 | spl66_94 | spl66_93 | spl66_92 | spl66_91 | spl66_90 | spl66_89 | spl66_88 | spl66_87 | spl66_86 | spl66_85 | spl66_84 | spl66_83 | spl66_82 | spl66_81 | spl66_80 | spl66_79 | spl66_78 | spl66_77 | spl66_76 | spl66_75 | spl66_74 | spl66_73 | spl66_72 | spl66_71 | spl66_70 | spl66_69 | spl66_68 | spl66_67 | spl66_66 | spl66_65 | ~ spl66_4), inference(avatar_split_clause, [], [f562, f602, f894, f902, f910, f918, f926, f934, f942, f950, f958, f966, f974, f982, f990, f998, f1006, f1014, f1022, f1030, f1038, f1046, f1054, f1062, f1070, f1078, f1086, f1094, f1102, f1110, f1118, f1126, f1134, f1142, f1150, f1158, f1166, f1174, f1182, f1190, f1198, f1206, f1214, f1222, f1230, f1238, f1246, f1254, f1262, f1270, f1278, f1286, f1294, f1302, f1310, f1318, f1326, f1334, f1342, f1350, f1358, f1366, f1374, f1382, f1390])).
fof(f562, plain, (~ (e3 = op(e3, e3)) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3), inference(cnf_transformation, [], [f120])).
fof(f1420, plain, (~ spl66_130 | spl66_61), inference(avatar_split_clause, [], [f551, f845, f1414])).
fof(f551, plain, ((e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3))) & ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2))) & ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1))) & ((e0 = op(e0, e0)) | ~ (e0 = op(e0, e0))) & (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f54])).
fof(f1419, plain, (~ spl66_130 | ~ spl66_41 | spl66_46), inference(avatar_split_clause, [], [f553, f781, f760, f1414])).
fof(f553, plain, ((e1 = op(e1, e0)) | ~ (e0 = op(e1, e1)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1418, plain, (~ spl66_130 | ~ spl66_21 | spl66_31), inference(avatar_split_clause, [], [f554, f717, f675, f1414])).
fof(f554, plain, ((e2 = op(e2, e0)) | ~ (e0 = op(e2, e2)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1417, plain, (~ spl66_130 | ~ spl66_1 | spl66_16), inference(avatar_split_clause, [], [f555, f653, f590, f1414])).
fof(f555, plain, ((e3 = op(e3, e0)) | ~ (e0 = op(e3, e3)) | ~ sP0), inference(cnf_transformation, [], [f186])).
fof(f1412, plain, (~ spl66_129 | spl66_42), inference(avatar_split_clause, [], [f546, f764, f1406])).
fof(f546, plain, ((e1 = op(e1, e1)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3))) & ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2))) & ((e1 = op(e1, e1)) | ~ (e1 = op(e1, e1))) & ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1)) & (e1 = op(e1, e1))) | ~ sP1), inference(nnf_transformation, [], [f55])).
fof(f1411, plain, (~ spl66_129 | ~ spl66_62 | spl66_57), inference(avatar_split_clause, [], [f547, f828, f849, f1406])).
fof(f547, plain, ((e0 = op(e0, e1)) | ~ (op(e0, e0) = e1) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1410, plain, (~ spl66_129 | ~ spl66_22 | spl66_27), inference(avatar_split_clause, [], [f549, f700, f679, f1406])).
fof(f549, plain, ((e2 = op(e2, e1)) | ~ (e1 = op(e2, e2)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1409, plain, (~ spl66_129 | ~ spl66_2 | spl66_12), inference(avatar_split_clause, [], [f550, f636, f594, f1406])).
fof(f550, plain, ((e3 = op(e3, e1)) | ~ (e1 = op(e3, e3)) | ~ sP1), inference(cnf_transformation, [], [f185])).
fof(f1404, plain, (~ spl66_128 | spl66_23), inference(avatar_split_clause, [], [f541, f683, f1398])).
fof(f541, plain, ((e2 = op(e2, e2)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3))) & ((e2 = op(e2, e2)) | ~ (e2 = op(e2, e2))) & ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1))) & ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2)) & (e2 = op(e2, e2))) | ~ sP2), inference(nnf_transformation, [], [f56])).
fof(f1403, plain, (~ spl66_128 | ~ spl66_63 | spl66_53), inference(avatar_split_clause, [], [f542, f811, f853, f1398])).
fof(f542, plain, ((e0 = op(e0, e2)) | ~ (op(e0, e0) = e2) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1402, plain, (~ spl66_128 | ~ spl66_43 | spl66_38), inference(avatar_split_clause, [], [f543, f747, f768, f1398])).
fof(f543, plain, ((e1 = op(e1, e2)) | ~ (e2 = op(e1, e1)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1401, plain, (~ spl66_128 | ~ spl66_3 | spl66_8), inference(avatar_split_clause, [], [f545, f619, f598, f1398])).
fof(f545, plain, ((e3 = op(e3, e2)) | ~ (e2 = op(e3, e3)) | ~ sP2), inference(cnf_transformation, [], [f184])).
fof(f1396, plain, (~ spl66_127 | spl66_61), inference(avatar_split_clause, [], [f537, f845, f1390])).
fof(f537, plain, ((e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP3), inference(nnf_transformation, [], [f57])).
fof(f1395, plain, (~ spl66_127 | ~ spl66_61), inference(avatar_split_clause, [], [f538, f845, f1390])).
fof(f538, plain, (~ (e0 = op(e0, e0)) | ~ sP3), inference(cnf_transformation, [], [f183])).
fof(f1388, plain, (~ spl66_126 | spl66_61), inference(avatar_split_clause, [], [f533, f845, f1382])).
fof(f533, plain, ((e0 = op(e0, e0)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP4), inference(nnf_transformation, [], [f58])).
fof(f1387, plain, (~ spl66_126 | ~ spl66_61), inference(avatar_split_clause, [], [f534, f845, f1382])).
fof(f534, plain, (~ (e0 = op(e0, e0)) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f1380, plain, (~ spl66_125 | spl66_61), inference(avatar_split_clause, [], [f529, f845, f1374])).
fof(f529, plain, ((e0 = op(e0, e0)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP5), inference(nnf_transformation, [], [f59])).
fof(f1379, plain, (~ spl66_125 | ~ spl66_61), inference(avatar_split_clause, [], [f530, f845, f1374])).
fof(f530, plain, (~ (e0 = op(e0, e0)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f1377, plain, (~ spl66_125 | ~ spl66_31), inference(avatar_split_clause, [], [f532, f717, f1374])).
fof(f532, plain, (~ (e2 = op(e2, e0)) | ~ sP5), inference(cnf_transformation, [], [f181])).
fof(f1372, plain, (~ spl66_124 | spl66_61), inference(avatar_split_clause, [], [f525, f845, f1366])).
fof(f525, plain, ((e0 = op(e0, e0)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e0, e0)) & (e0 = op(e0, e0))) | ~ sP6), inference(nnf_transformation, [], [f60])).
fof(f1371, plain, (~ spl66_124 | ~ spl66_61), inference(avatar_split_clause, [], [f526, f845, f1366])).
fof(f526, plain, (~ (e0 = op(e0, e0)) | ~ sP6), inference(cnf_transformation, [], [f180])).
fof(f1363, plain, (~ spl66_123 | ~ spl66_62), inference(avatar_split_clause, [], [f522, f849, f1358])).
fof(f522, plain, (~ (op(e0, e0) = e1) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP7), inference(nnf_transformation, [], [f61])).
fof(f1362, plain, (~ spl66_123 | spl66_62), inference(avatar_split_clause, [], [f523, f849, f1358])).
fof(f523, plain, ((op(e0, e0) = e1) | ~ sP7), inference(cnf_transformation, [], [f179])).
fof(f1354, plain, (~ spl66_122 | spl66_42), inference(avatar_split_clause, [], [f519, f764, f1350])).
fof(f519, plain, ((e1 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP8), inference(nnf_transformation, [], [f62])).
fof(f1353, plain, (~ spl66_122 | ~ spl66_42), inference(avatar_split_clause, [], [f520, f764, f1350])).
fof(f520, plain, (~ (e1 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f178])).
fof(f1348, plain, (~ spl66_121 | spl66_57), inference(avatar_split_clause, [], [f513, f828, f1342])).
fof(f513, plain, ((e0 = op(e0, e1)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP9), inference(nnf_transformation, [], [f63])).
fof(f1347, plain, (~ spl66_121 | ~ spl66_62), inference(avatar_split_clause, [], [f514, f849, f1342])).
fof(f514, plain, (~ (op(e0, e0) = e1) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1346, plain, (~ spl66_121 | spl66_22), inference(avatar_split_clause, [], [f515, f679, f1342])).
fof(f515, plain, ((e1 = op(e2, e2)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1345, plain, (~ spl66_121 | ~ spl66_27), inference(avatar_split_clause, [], [f516, f700, f1342])).
fof(f516, plain, (~ (e2 = op(e2, e1)) | ~ sP9), inference(cnf_transformation, [], [f177])).
fof(f1340, plain, (~ spl66_120 | spl66_57), inference(avatar_split_clause, [], [f509, f828, f1334])).
fof(f509, plain, ((e0 = op(e0, e1)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (op(e0, e0) = e1) & (e0 = op(e0, e1))) | ~ sP10), inference(nnf_transformation, [], [f64])).
fof(f1339, plain, (~ spl66_120 | ~ spl66_62), inference(avatar_split_clause, [], [f510, f849, f1334])).
fof(f510, plain, (~ (op(e0, e0) = e1) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1338, plain, (~ spl66_120 | spl66_2), inference(avatar_split_clause, [], [f511, f594, f1334])).
fof(f511, plain, ((e1 = op(e3, e3)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1337, plain, (~ spl66_120 | ~ spl66_12), inference(avatar_split_clause, [], [f512, f636, f1334])).
fof(f512, plain, (~ (e3 = op(e3, e1)) | ~ sP10), inference(cnf_transformation, [], [f176])).
fof(f1331, plain, (~ spl66_119 | ~ spl66_63), inference(avatar_split_clause, [], [f506, f853, f1326])).
fof(f506, plain, (~ (op(e0, e0) = e2) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP11), inference(nnf_transformation, [], [f65])).
fof(f1330, plain, (~ spl66_119 | spl66_63), inference(avatar_split_clause, [], [f507, f853, f1326])).
fof(f507, plain, ((op(e0, e0) = e2) | ~ sP11), inference(cnf_transformation, [], [f175])).
fof(f1324, plain, (~ spl66_118 | spl66_53), inference(avatar_split_clause, [], [f501, f811, f1318])).
fof(f501, plain, ((e0 = op(e0, e2)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP12), inference(nnf_transformation, [], [f66])).
fof(f1323, plain, (~ spl66_118 | ~ spl66_63), inference(avatar_split_clause, [], [f502, f853, f1318])).
fof(f502, plain, (~ (op(e0, e0) = e2) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1322, plain, (~ spl66_118 | spl66_43), inference(avatar_split_clause, [], [f503, f768, f1318])).
fof(f503, plain, ((e2 = op(e1, e1)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1321, plain, (~ spl66_118 | ~ spl66_38), inference(avatar_split_clause, [], [f504, f747, f1318])).
fof(f504, plain, (~ (e1 = op(e1, e2)) | ~ sP12), inference(cnf_transformation, [], [f174])).
fof(f1316, plain, (~ spl66_117 | spl66_53), inference(avatar_split_clause, [], [f497, f811, f1310])).
fof(f497, plain, ((e0 = op(e0, e2)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP13), inference(nnf_transformation, [], [f67])).
fof(f1314, plain, (~ spl66_117 | spl66_23), inference(avatar_split_clause, [], [f499, f683, f1310])).
fof(f499, plain, ((e2 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1313, plain, (~ spl66_117 | ~ spl66_23), inference(avatar_split_clause, [], [f500, f683, f1310])).
fof(f500, plain, (~ (e2 = op(e2, e2)) | ~ sP13), inference(cnf_transformation, [], [f173])).
fof(f1308, plain, (~ spl66_116 | spl66_53), inference(avatar_split_clause, [], [f493, f811, f1302])).
fof(f493, plain, ((e0 = op(e0, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (op(e0, e0) = e2) & (e0 = op(e0, e2))) | ~ sP14), inference(nnf_transformation, [], [f68])).
fof(f1307, plain, (~ spl66_116 | ~ spl66_63), inference(avatar_split_clause, [], [f494, f853, f1302])).
fof(f494, plain, (~ (op(e0, e0) = e2) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1306, plain, (~ spl66_116 | spl66_3), inference(avatar_split_clause, [], [f495, f598, f1302])).
fof(f495, plain, ((e2 = op(e3, e3)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1305, plain, (~ spl66_116 | ~ spl66_8), inference(avatar_split_clause, [], [f496, f619, f1302])).
fof(f496, plain, (~ (e3 = op(e3, e2)) | ~ sP14), inference(cnf_transformation, [], [f172])).
fof(f1299, plain, (~ spl66_115 | ~ spl66_64), inference(avatar_split_clause, [], [f490, f857, f1294])).
fof(f490, plain, (~ (op(e0, e0) = e3) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP15), inference(nnf_transformation, [], [f69])).
fof(f1298, plain, (~ spl66_115 | spl66_64), inference(avatar_split_clause, [], [f491, f857, f1294])).
fof(f491, plain, ((op(e0, e0) = e3) | ~ sP15), inference(cnf_transformation, [], [f171])).
fof(f1292, plain, (~ spl66_114 | spl66_49), inference(avatar_split_clause, [], [f485, f794, f1286])).
fof(f485, plain, ((e0 = op(e0, e3)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP16), inference(nnf_transformation, [], [f70])).
fof(f1291, plain, (~ spl66_114 | ~ spl66_64), inference(avatar_split_clause, [], [f486, f857, f1286])).
fof(f486, plain, (~ (op(e0, e0) = e3) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1290, plain, (~ spl66_114 | spl66_44), inference(avatar_split_clause, [], [f487, f772, f1286])).
fof(f487, plain, ((e3 = op(e1, e1)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1289, plain, (~ spl66_114 | ~ spl66_34), inference(avatar_split_clause, [], [f488, f730, f1286])).
fof(f488, plain, (~ (e1 = op(e1, e3)) | ~ sP16), inference(cnf_transformation, [], [f170])).
fof(f1284, plain, (~ spl66_113 | spl66_49), inference(avatar_split_clause, [], [f481, f794, f1278])).
fof(f481, plain, ((e0 = op(e0, e3)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP17), inference(nnf_transformation, [], [f71])).
fof(f1283, plain, (~ spl66_113 | ~ spl66_64), inference(avatar_split_clause, [], [f482, f857, f1278])).
fof(f482, plain, (~ (op(e0, e0) = e3) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1282, plain, (~ spl66_113 | spl66_24), inference(avatar_split_clause, [], [f483, f687, f1278])).
fof(f483, plain, ((e3 = op(e2, e2)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1281, plain, (~ spl66_113 | ~ spl66_19), inference(avatar_split_clause, [], [f484, f666, f1278])).
fof(f484, plain, (~ (e2 = op(e2, e3)) | ~ sP17), inference(cnf_transformation, [], [f169])).
fof(f1276, plain, (~ spl66_112 | spl66_49), inference(avatar_split_clause, [], [f477, f794, f1270])).
fof(f477, plain, ((e0 = op(e0, e3)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (op(e0, e0) = e3) & (e0 = op(e0, e3))) | ~ sP18), inference(nnf_transformation, [], [f72])).
fof(f1274, plain, (~ spl66_112 | spl66_4), inference(avatar_split_clause, [], [f479, f602, f1270])).
fof(f479, plain, ((e3 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1273, plain, (~ spl66_112 | ~ spl66_4), inference(avatar_split_clause, [], [f480, f602, f1270])).
fof(f480, plain, (~ (e3 = op(e3, e3)) | ~ sP18), inference(cnf_transformation, [], [f168])).
fof(f1266, plain, (~ spl66_111 | spl66_61), inference(avatar_split_clause, [], [f475, f845, f1262])).
fof(f475, plain, ((e0 = op(e0, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP19), inference(nnf_transformation, [], [f73])).
fof(f1265, plain, (~ spl66_111 | ~ spl66_61), inference(avatar_split_clause, [], [f476, f845, f1262])).
fof(f476, plain, (~ (e0 = op(e0, e0)) | ~ sP19), inference(cnf_transformation, [], [f167])).
fof(f1259, plain, (~ spl66_110 | ~ spl66_41), inference(avatar_split_clause, [], [f470, f760, f1254])).
fof(f470, plain, (~ (e0 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP20), inference(nnf_transformation, [], [f74])).
fof(f1258, plain, (~ spl66_110 | spl66_41), inference(avatar_split_clause, [], [f471, f760, f1254])).
fof(f471, plain, ((e0 = op(e1, e1)) | ~ sP20), inference(cnf_transformation, [], [f166])).
fof(f1252, plain, (~ spl66_109 | spl66_46), inference(avatar_split_clause, [], [f465, f781, f1246])).
fof(f465, plain, ((e1 = op(e1, e0)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP21), inference(nnf_transformation, [], [f75])).
fof(f1251, plain, (~ spl66_109 | ~ spl66_41), inference(avatar_split_clause, [], [f466, f760, f1246])).
fof(f466, plain, (~ (e0 = op(e1, e1)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1250, plain, (~ spl66_109 | spl66_21), inference(avatar_split_clause, [], [f467, f675, f1246])).
fof(f467, plain, ((e0 = op(e2, e2)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1249, plain, (~ spl66_109 | ~ spl66_31), inference(avatar_split_clause, [], [f468, f717, f1246])).
fof(f468, plain, (~ (e2 = op(e2, e0)) | ~ sP21), inference(cnf_transformation, [], [f165])).
fof(f1244, plain, (~ spl66_108 | spl66_46), inference(avatar_split_clause, [], [f461, f781, f1238])).
fof(f461, plain, ((e1 = op(e1, e0)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e1, e1)) & (e1 = op(e1, e0))) | ~ sP22), inference(nnf_transformation, [], [f76])).
fof(f1243, plain, (~ spl66_108 | ~ spl66_41), inference(avatar_split_clause, [], [f462, f760, f1238])).
fof(f462, plain, (~ (e0 = op(e1, e1)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1242, plain, (~ spl66_108 | spl66_1), inference(avatar_split_clause, [], [f463, f590, f1238])).
fof(f463, plain, ((e0 = op(e3, e3)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1241, plain, (~ spl66_108 | ~ spl66_16), inference(avatar_split_clause, [], [f464, f653, f1238])).
fof(f464, plain, (~ (e3 = op(e3, e0)) | ~ sP22), inference(cnf_transformation, [], [f164])).
fof(f1236, plain, (~ spl66_107 | spl66_42), inference(avatar_split_clause, [], [f457, f764, f1230])).
fof(f457, plain, ((e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP23), inference(nnf_transformation, [], [f77])).
fof(f1235, plain, (~ spl66_107 | ~ spl66_42), inference(avatar_split_clause, [], [f458, f764, f1230])).
fof(f458, plain, (~ (e1 = op(e1, e1)) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f1234, plain, (~ spl66_107 | spl66_62), inference(avatar_split_clause, [], [f459, f849, f1230])).
fof(f459, plain, ((op(e0, e0) = e1) | ~ sP23), inference(cnf_transformation, [], [f163])).
fof(f1228, plain, (~ spl66_106 | spl66_42), inference(avatar_split_clause, [], [f453, f764, f1222])).
fof(f453, plain, ((e1 = op(e1, e1)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP24), inference(nnf_transformation, [], [f78])).
fof(f1227, plain, (~ spl66_106 | ~ spl66_42), inference(avatar_split_clause, [], [f454, f764, f1222])).
fof(f454, plain, (~ (e1 = op(e1, e1)) | ~ sP24), inference(cnf_transformation, [], [f162])).
fof(f1220, plain, (~ spl66_105 | spl66_42), inference(avatar_split_clause, [], [f449, f764, f1214])).
fof(f449, plain, ((e1 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP25), inference(nnf_transformation, [], [f79])).
fof(f1219, plain, (~ spl66_105 | ~ spl66_42), inference(avatar_split_clause, [], [f450, f764, f1214])).
fof(f450, plain, (~ (e1 = op(e1, e1)) | ~ sP25), inference(cnf_transformation, [], [f161])).
fof(f1212, plain, (~ spl66_104 | spl66_42), inference(avatar_split_clause, [], [f445, f764, f1206])).
fof(f445, plain, ((e1 = op(e1, e1)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e1, e1)) & (e1 = op(e1, e1))) | ~ sP26), inference(nnf_transformation, [], [f80])).
fof(f1211, plain, (~ spl66_104 | ~ spl66_42), inference(avatar_split_clause, [], [f446, f764, f1206])).
fof(f446, plain, (~ (e1 = op(e1, e1)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1210, plain, (~ spl66_104 | spl66_2), inference(avatar_split_clause, [], [f447, f594, f1206])).
fof(f447, plain, ((e1 = op(e3, e3)) | ~ sP26), inference(cnf_transformation, [], [f160])).
fof(f1204, plain, (~ spl66_103 | spl66_38), inference(avatar_split_clause, [], [f441, f747, f1198])).
fof(f441, plain, ((e1 = op(e1, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP27), inference(nnf_transformation, [], [f81])).
fof(f1203, plain, (~ spl66_103 | ~ spl66_43), inference(avatar_split_clause, [], [f442, f768, f1198])).
fof(f442, plain, (~ (e2 = op(e1, e1)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1202, plain, (~ spl66_103 | spl66_63), inference(avatar_split_clause, [], [f443, f853, f1198])).
fof(f443, plain, ((op(e0, e0) = e2) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1201, plain, (~ spl66_103 | ~ spl66_53), inference(avatar_split_clause, [], [f444, f811, f1198])).
fof(f444, plain, (~ (e0 = op(e0, e2)) | ~ sP27), inference(cnf_transformation, [], [f159])).
fof(f1196, plain, (~ spl66_102 | spl66_38), inference(avatar_split_clause, [], [f437, f747, f1190])).
fof(f437, plain, ((e1 = op(e1, e2)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP28), inference(nnf_transformation, [], [f82])).
fof(f1195, plain, (~ spl66_102 | ~ spl66_43), inference(avatar_split_clause, [], [f438, f768, f1190])).
fof(f438, plain, (~ (e2 = op(e1, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1194, plain, (~ spl66_102 | spl66_43), inference(avatar_split_clause, [], [f439, f768, f1190])).
fof(f439, plain, ((e2 = op(e1, e1)) | ~ sP28), inference(cnf_transformation, [], [f158])).
fof(f1186, plain, (~ spl66_101 | spl66_23), inference(avatar_split_clause, [], [f435, f683, f1182])).
fof(f435, plain, ((e2 = op(e2, e2)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP29), inference(nnf_transformation, [], [f83])).
fof(f1185, plain, (~ spl66_101 | ~ spl66_23), inference(avatar_split_clause, [], [f436, f683, f1182])).
fof(f436, plain, (~ (e2 = op(e2, e2)) | ~ sP29), inference(cnf_transformation, [], [f157])).
fof(f1180, plain, (~ spl66_100 | spl66_38), inference(avatar_split_clause, [], [f429, f747, f1174])).
fof(f429, plain, ((e1 = op(e1, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e1, e1)) & (e1 = op(e1, e2))) | ~ sP30), inference(nnf_transformation, [], [f84])).
fof(f1179, plain, (~ spl66_100 | ~ spl66_43), inference(avatar_split_clause, [], [f430, f768, f1174])).
fof(f430, plain, (~ (e2 = op(e1, e1)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1178, plain, (~ spl66_100 | spl66_3), inference(avatar_split_clause, [], [f431, f598, f1174])).
fof(f431, plain, ((e2 = op(e3, e3)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1177, plain, (~ spl66_100 | ~ spl66_8), inference(avatar_split_clause, [], [f432, f619, f1174])).
fof(f432, plain, (~ (e3 = op(e3, e2)) | ~ sP30), inference(cnf_transformation, [], [f156])).
fof(f1172, plain, (~ spl66_99 | spl66_34), inference(avatar_split_clause, [], [f425, f730, f1166])).
fof(f425, plain, ((e1 = op(e1, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP31), inference(nnf_transformation, [], [f85])).
fof(f1171, plain, (~ spl66_99 | ~ spl66_44), inference(avatar_split_clause, [], [f426, f772, f1166])).
fof(f426, plain, (~ (e3 = op(e1, e1)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1170, plain, (~ spl66_99 | spl66_64), inference(avatar_split_clause, [], [f427, f857, f1166])).
fof(f427, plain, ((op(e0, e0) = e3) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1169, plain, (~ spl66_99 | ~ spl66_49), inference(avatar_split_clause, [], [f428, f794, f1166])).
fof(f428, plain, (~ (e0 = op(e0, e3)) | ~ sP31), inference(cnf_transformation, [], [f155])).
fof(f1163, plain, (~ spl66_98 | ~ spl66_44), inference(avatar_split_clause, [], [f422, f772, f1158])).
fof(f422, plain, (~ (e3 = op(e1, e1)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP32), inference(nnf_transformation, [], [f86])).
fof(f1162, plain, (~ spl66_98 | spl66_44), inference(avatar_split_clause, [], [f423, f772, f1158])).
fof(f423, plain, ((e3 = op(e1, e1)) | ~ sP32), inference(cnf_transformation, [], [f154])).
fof(f1156, plain, (~ spl66_97 | spl66_34), inference(avatar_split_clause, [], [f417, f730, f1150])).
fof(f417, plain, ((e1 = op(e1, e3)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP33), inference(nnf_transformation, [], [f87])).
fof(f1155, plain, (~ spl66_97 | ~ spl66_44), inference(avatar_split_clause, [], [f418, f772, f1150])).
fof(f418, plain, (~ (e3 = op(e1, e1)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1154, plain, (~ spl66_97 | spl66_24), inference(avatar_split_clause, [], [f419, f687, f1150])).
fof(f419, plain, ((e3 = op(e2, e2)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1153, plain, (~ spl66_97 | ~ spl66_19), inference(avatar_split_clause, [], [f420, f666, f1150])).
fof(f420, plain, (~ (e2 = op(e2, e3)) | ~ sP33), inference(cnf_transformation, [], [f153])).
fof(f1146, plain, (~ spl66_96 | spl66_4), inference(avatar_split_clause, [], [f415, f602, f1142])).
fof(f415, plain, ((e3 = op(e3, e3)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e1, e1)) & (e1 = op(e1, e3))) | ~ sP34), inference(nnf_transformation, [], [f88])).
fof(f1145, plain, (~ spl66_96 | ~ spl66_4), inference(avatar_split_clause, [], [f416, f602, f1142])).
fof(f416, plain, (~ (e3 = op(e3, e3)) | ~ sP34), inference(cnf_transformation, [], [f152])).
fof(f1138, plain, (~ spl66_95 | spl66_61), inference(avatar_split_clause, [], [f411, f845, f1134])).
fof(f411, plain, ((e0 = op(e0, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP35), inference(nnf_transformation, [], [f89])).
fof(f1137, plain, (~ spl66_95 | ~ spl66_61), inference(avatar_split_clause, [], [f412, f845, f1134])).
fof(f412, plain, (~ (e0 = op(e0, e0)) | ~ sP35), inference(cnf_transformation, [], [f151])).
fof(f1132, plain, (~ spl66_94 | spl66_31), inference(avatar_split_clause, [], [f405, f717, f1126])).
fof(f405, plain, ((e2 = op(e2, e0)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP36), inference(nnf_transformation, [], [f90])).
fof(f1131, plain, (~ spl66_94 | ~ spl66_21), inference(avatar_split_clause, [], [f406, f675, f1126])).
fof(f406, plain, (~ (e0 = op(e2, e2)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1130, plain, (~ spl66_94 | spl66_41), inference(avatar_split_clause, [], [f407, f760, f1126])).
fof(f407, plain, ((e0 = op(e1, e1)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1129, plain, (~ spl66_94 | ~ spl66_46), inference(avatar_split_clause, [], [f408, f781, f1126])).
fof(f408, plain, (~ (e1 = op(e1, e0)) | ~ sP36), inference(cnf_transformation, [], [f150])).
fof(f1124, plain, (~ spl66_93 | spl66_31), inference(avatar_split_clause, [], [f401, f717, f1118])).
fof(f401, plain, ((e2 = op(e2, e0)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP37), inference(nnf_transformation, [], [f91])).
fof(f1123, plain, (~ spl66_93 | ~ spl66_21), inference(avatar_split_clause, [], [f402, f675, f1118])).
fof(f402, plain, (~ (e0 = op(e2, e2)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1122, plain, (~ spl66_93 | spl66_21), inference(avatar_split_clause, [], [f403, f675, f1118])).
fof(f403, plain, ((e0 = op(e2, e2)) | ~ sP37), inference(cnf_transformation, [], [f149])).
fof(f1116, plain, (~ spl66_92 | spl66_31), inference(avatar_split_clause, [], [f397, f717, f1110])).
fof(f397, plain, ((e2 = op(e2, e0)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e2, e2)) & (e2 = op(e2, e0))) | ~ sP38), inference(nnf_transformation, [], [f92])).
fof(f1115, plain, (~ spl66_92 | ~ spl66_21), inference(avatar_split_clause, [], [f398, f675, f1110])).
fof(f398, plain, (~ (e0 = op(e2, e2)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1114, plain, (~ spl66_92 | spl66_1), inference(avatar_split_clause, [], [f399, f590, f1110])).
fof(f399, plain, ((e0 = op(e3, e3)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1113, plain, (~ spl66_92 | ~ spl66_16), inference(avatar_split_clause, [], [f400, f653, f1110])).
fof(f400, plain, (~ (e3 = op(e3, e0)) | ~ sP38), inference(cnf_transformation, [], [f148])).
fof(f1108, plain, (~ spl66_91 | spl66_27), inference(avatar_split_clause, [], [f393, f700, f1102])).
fof(f393, plain, ((e2 = op(e2, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP39), inference(nnf_transformation, [], [f93])).
fof(f1107, plain, (~ spl66_91 | ~ spl66_22), inference(avatar_split_clause, [], [f394, f679, f1102])).
fof(f394, plain, (~ (e1 = op(e2, e2)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1106, plain, (~ spl66_91 | spl66_62), inference(avatar_split_clause, [], [f395, f849, f1102])).
fof(f395, plain, ((op(e0, e0) = e1) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1105, plain, (~ spl66_91 | ~ spl66_57), inference(avatar_split_clause, [], [f396, f828, f1102])).
fof(f396, plain, (~ (e0 = op(e0, e1)) | ~ sP39), inference(cnf_transformation, [], [f147])).
fof(f1098, plain, (~ spl66_90 | spl66_42), inference(avatar_split_clause, [], [f391, f764, f1094])).
fof(f391, plain, ((e1 = op(e1, e1)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP40), inference(nnf_transformation, [], [f94])).
fof(f1097, plain, (~ spl66_90 | ~ spl66_42), inference(avatar_split_clause, [], [f392, f764, f1094])).
fof(f392, plain, (~ (e1 = op(e1, e1)) | ~ sP40), inference(cnf_transformation, [], [f146])).
fof(f1092, plain, (~ spl66_89 | spl66_27), inference(avatar_split_clause, [], [f385, f700, f1086])).
fof(f385, plain, ((e2 = op(e2, e1)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP41), inference(nnf_transformation, [], [f95])).
fof(f1091, plain, (~ spl66_89 | ~ spl66_22), inference(avatar_split_clause, [], [f386, f679, f1086])).
fof(f386, plain, (~ (e1 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1090, plain, (~ spl66_89 | spl66_22), inference(avatar_split_clause, [], [f387, f679, f1086])).
fof(f387, plain, ((e1 = op(e2, e2)) | ~ sP41), inference(cnf_transformation, [], [f145])).
fof(f1084, plain, (~ spl66_88 | spl66_27), inference(avatar_split_clause, [], [f381, f700, f1078])).
fof(f381, plain, ((e2 = op(e2, e1)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e2, e2)) & (e2 = op(e2, e1))) | ~ sP42), inference(nnf_transformation, [], [f96])).
fof(f1083, plain, (~ spl66_88 | ~ spl66_22), inference(avatar_split_clause, [], [f382, f679, f1078])).
fof(f382, plain, (~ (e1 = op(e2, e2)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1082, plain, (~ spl66_88 | spl66_2), inference(avatar_split_clause, [], [f383, f594, f1078])).
fof(f383, plain, ((e1 = op(e3, e3)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1081, plain, (~ spl66_88 | ~ spl66_12), inference(avatar_split_clause, [], [f384, f636, f1078])).
fof(f384, plain, (~ (e3 = op(e3, e1)) | ~ sP42), inference(cnf_transformation, [], [f144])).
fof(f1076, plain, (~ spl66_87 | spl66_23), inference(avatar_split_clause, [], [f377, f683, f1070])).
fof(f377, plain, ((e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP43), inference(nnf_transformation, [], [f97])).
fof(f1075, plain, (~ spl66_87 | ~ spl66_23), inference(avatar_split_clause, [], [f378, f683, f1070])).
fof(f378, plain, (~ (e2 = op(e2, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1074, plain, (~ spl66_87 | spl66_63), inference(avatar_split_clause, [], [f379, f853, f1070])).
fof(f379, plain, ((op(e0, e0) = e2) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1073, plain, (~ spl66_87 | ~ spl66_53), inference(avatar_split_clause, [], [f380, f811, f1070])).
fof(f380, plain, (~ (e0 = op(e0, e2)) | ~ sP43), inference(cnf_transformation, [], [f143])).
fof(f1068, plain, (~ spl66_86 | spl66_23), inference(avatar_split_clause, [], [f373, f683, f1062])).
fof(f373, plain, ((e2 = op(e2, e2)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP44), inference(nnf_transformation, [], [f98])).
fof(f1067, plain, (~ spl66_86 | ~ spl66_23), inference(avatar_split_clause, [], [f374, f683, f1062])).
fof(f374, plain, (~ (e2 = op(e2, e2)) | ~ sP44), inference(cnf_transformation, [], [f142])).
fof(f1060, plain, (~ spl66_85 | spl66_23), inference(avatar_split_clause, [], [f369, f683, f1054])).
fof(f369, plain, ((e2 = op(e2, e2)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP45), inference(nnf_transformation, [], [f99])).
fof(f1059, plain, (~ spl66_85 | ~ spl66_23), inference(avatar_split_clause, [], [f370, f683, f1054])).
fof(f370, plain, (~ (e2 = op(e2, e2)) | ~ sP45), inference(cnf_transformation, [], [f141])).
fof(f1052, plain, (~ spl66_84 | spl66_23), inference(avatar_split_clause, [], [f365, f683, f1046])).
fof(f365, plain, ((e2 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & (e2 = op(e2, e2))) | ~ sP46), inference(nnf_transformation, [], [f100])).
fof(f1051, plain, (~ spl66_84 | ~ spl66_23), inference(avatar_split_clause, [], [f366, f683, f1046])).
fof(f366, plain, (~ (e2 = op(e2, e2)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f1050, plain, (~ spl66_84 | spl66_3), inference(avatar_split_clause, [], [f367, f598, f1046])).
fof(f367, plain, ((e2 = op(e3, e3)) | ~ sP46), inference(cnf_transformation, [], [f140])).
fof(f1044, plain, (~ spl66_83 | spl66_19), inference(avatar_split_clause, [], [f361, f666, f1038])).
fof(f361, plain, ((e2 = op(e2, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP47), inference(nnf_transformation, [], [f101])).
fof(f1043, plain, (~ spl66_83 | ~ spl66_24), inference(avatar_split_clause, [], [f362, f687, f1038])).
fof(f362, plain, (~ (e3 = op(e2, e2)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1042, plain, (~ spl66_83 | spl66_64), inference(avatar_split_clause, [], [f363, f857, f1038])).
fof(f363, plain, ((op(e0, e0) = e3) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1041, plain, (~ spl66_83 | ~ spl66_49), inference(avatar_split_clause, [], [f364, f794, f1038])).
fof(f364, plain, (~ (e0 = op(e0, e3)) | ~ sP47), inference(cnf_transformation, [], [f139])).
fof(f1036, plain, (~ spl66_82 | spl66_19), inference(avatar_split_clause, [], [f357, f666, f1030])).
fof(f357, plain, ((e2 = op(e2, e3)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP48), inference(nnf_transformation, [], [f102])).
fof(f1035, plain, (~ spl66_82 | ~ spl66_24), inference(avatar_split_clause, [], [f358, f687, f1030])).
fof(f358, plain, (~ (e3 = op(e2, e2)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1034, plain, (~ spl66_82 | spl66_44), inference(avatar_split_clause, [], [f359, f772, f1030])).
fof(f359, plain, ((e3 = op(e1, e1)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1033, plain, (~ spl66_82 | ~ spl66_34), inference(avatar_split_clause, [], [f360, f730, f1030])).
fof(f360, plain, (~ (e1 = op(e1, e3)) | ~ sP48), inference(cnf_transformation, [], [f138])).
fof(f1027, plain, (~ spl66_81 | ~ spl66_24), inference(avatar_split_clause, [], [f354, f687, f1022])).
fof(f354, plain, (~ (e3 = op(e2, e2)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP49), inference(nnf_transformation, [], [f103])).
fof(f1026, plain, (~ spl66_81 | spl66_24), inference(avatar_split_clause, [], [f355, f687, f1022])).
fof(f355, plain, ((e3 = op(e2, e2)) | ~ sP49), inference(cnf_transformation, [], [f137])).
fof(f1020, plain, (~ spl66_80 | spl66_19), inference(avatar_split_clause, [], [f349, f666, f1014])).
fof(f349, plain, ((e2 = op(e2, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e3 = op(e3, e3)) & (e3 = op(e3, e3)) & ~ (e3 = op(e2, e2)) & (e2 = op(e2, e3))) | ~ sP50), inference(nnf_transformation, [], [f104])).
fof(f1018, plain, (~ spl66_80 | spl66_4), inference(avatar_split_clause, [], [f351, f602, f1014])).
fof(f351, plain, ((e3 = op(e3, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1017, plain, (~ spl66_80 | ~ spl66_4), inference(avatar_split_clause, [], [f352, f602, f1014])).
fof(f352, plain, (~ (e3 = op(e3, e3)) | ~ sP50), inference(cnf_transformation, [], [f136])).
fof(f1010, plain, (~ spl66_79 | spl66_61), inference(avatar_split_clause, [], [f347, f845, f1006])).
fof(f347, plain, ((e0 = op(e0, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e0 = op(e0, e0)) & (e0 = op(e0, e0)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP51), inference(nnf_transformation, [], [f105])).
fof(f1009, plain, (~ spl66_79 | ~ spl66_61), inference(avatar_split_clause, [], [f348, f845, f1006])).
fof(f348, plain, (~ (e0 = op(e0, e0)) | ~ sP51), inference(cnf_transformation, [], [f135])).
fof(f1004, plain, (~ spl66_78 | spl66_16), inference(avatar_split_clause, [], [f341, f653, f998])).
fof(f341, plain, ((e3 = op(e3, e0)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e1 = op(e1, e0)) & (e0 = op(e1, e1)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP52), inference(nnf_transformation, [], [f106])).
fof(f1003, plain, (~ spl66_78 | ~ spl66_1), inference(avatar_split_clause, [], [f342, f590, f998])).
fof(f342, plain, (~ (e0 = op(e3, e3)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1002, plain, (~ spl66_78 | spl66_41), inference(avatar_split_clause, [], [f343, f760, f998])).
fof(f343, plain, ((e0 = op(e1, e1)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f1001, plain, (~ spl66_78 | ~ spl66_46), inference(avatar_split_clause, [], [f344, f781, f998])).
fof(f344, plain, (~ (e1 = op(e1, e0)) | ~ sP52), inference(cnf_transformation, [], [f134])).
fof(f996, plain, (~ spl66_77 | spl66_16), inference(avatar_split_clause, [], [f337, f653, f990])).
fof(f337, plain, ((e3 = op(e3, e0)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e2 = op(e2, e0)) & (e0 = op(e2, e2)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP53), inference(nnf_transformation, [], [f107])).
fof(f995, plain, (~ spl66_77 | ~ spl66_1), inference(avatar_split_clause, [], [f338, f590, f990])).
fof(f338, plain, (~ (e0 = op(e3, e3)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f994, plain, (~ spl66_77 | spl66_21), inference(avatar_split_clause, [], [f339, f675, f990])).
fof(f339, plain, ((e0 = op(e2, e2)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f993, plain, (~ spl66_77 | ~ spl66_31), inference(avatar_split_clause, [], [f340, f717, f990])).
fof(f340, plain, (~ (e2 = op(e2, e0)) | ~ sP53), inference(cnf_transformation, [], [f133])).
fof(f988, plain, (~ spl66_76 | spl66_16), inference(avatar_split_clause, [], [f333, f653, f982])).
fof(f333, plain, ((e3 = op(e3, e0)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e3 = op(e3, e0)) & (e0 = op(e3, e3)) & ~ (e0 = op(e3, e3)) & (e3 = op(e3, e0))) | ~ sP54), inference(nnf_transformation, [], [f108])).
fof(f987, plain, (~ spl66_76 | ~ spl66_1), inference(avatar_split_clause, [], [f334, f590, f982])).
fof(f334, plain, (~ (e0 = op(e3, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f986, plain, (~ spl66_76 | spl66_1), inference(avatar_split_clause, [], [f335, f590, f982])).
fof(f335, plain, ((e0 = op(e3, e3)) | ~ sP54), inference(cnf_transformation, [], [f132])).
fof(f980, plain, (~ spl66_75 | spl66_12), inference(avatar_split_clause, [], [f329, f636, f974])).
fof(f329, plain, ((e3 = op(e3, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e0 = op(e0, e1)) & (op(e0, e0) = e1) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP55), inference(nnf_transformation, [], [f109])).
fof(f979, plain, (~ spl66_75 | ~ spl66_2), inference(avatar_split_clause, [], [f330, f594, f974])).
fof(f330, plain, (~ (e1 = op(e3, e3)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f978, plain, (~ spl66_75 | spl66_62), inference(avatar_split_clause, [], [f331, f849, f974])).
fof(f331, plain, ((op(e0, e0) = e1) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f977, plain, (~ spl66_75 | ~ spl66_57), inference(avatar_split_clause, [], [f332, f828, f974])).
fof(f332, plain, (~ (e0 = op(e0, e1)) | ~ sP55), inference(cnf_transformation, [], [f131])).
fof(f972, plain, (~ spl66_74 | spl66_12), inference(avatar_split_clause, [], [f325, f636, f966])).
fof(f325, plain, ((e3 = op(e3, e1)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e1 = op(e1, e1)) & (e1 = op(e1, e1)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP56), inference(nnf_transformation, [], [f110])).
fof(f970, plain, (~ spl66_74 | spl66_42), inference(avatar_split_clause, [], [f327, f764, f966])).
fof(f327, plain, ((e1 = op(e1, e1)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f969, plain, (~ spl66_74 | ~ spl66_42), inference(avatar_split_clause, [], [f328, f764, f966])).
fof(f328, plain, (~ (e1 = op(e1, e1)) | ~ sP56), inference(cnf_transformation, [], [f130])).
fof(f964, plain, (~ spl66_73 | spl66_12), inference(avatar_split_clause, [], [f321, f636, f958])).
fof(f321, plain, ((e3 = op(e3, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e2 = op(e2, e1)) & (e1 = op(e2, e2)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP57), inference(nnf_transformation, [], [f111])).
fof(f963, plain, (~ spl66_73 | ~ spl66_2), inference(avatar_split_clause, [], [f322, f594, f958])).
fof(f322, plain, (~ (e1 = op(e3, e3)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f962, plain, (~ spl66_73 | spl66_22), inference(avatar_split_clause, [], [f323, f679, f958])).
fof(f323, plain, ((e1 = op(e2, e2)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f961, plain, (~ spl66_73 | ~ spl66_27), inference(avatar_split_clause, [], [f324, f700, f958])).
fof(f324, plain, (~ (e2 = op(e2, e1)) | ~ sP57), inference(cnf_transformation, [], [f129])).
fof(f956, plain, (~ spl66_72 | spl66_12), inference(avatar_split_clause, [], [f317, f636, f950])).
fof(f317, plain, ((e3 = op(e3, e1)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e3 = op(e3, e1)) & (e1 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & (e3 = op(e3, e1))) | ~ sP58), inference(nnf_transformation, [], [f112])).
fof(f955, plain, (~ spl66_72 | ~ spl66_2), inference(avatar_split_clause, [], [f318, f594, f950])).
fof(f318, plain, (~ (e1 = op(e3, e3)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f954, plain, (~ spl66_72 | spl66_2), inference(avatar_split_clause, [], [f319, f594, f950])).
fof(f319, plain, ((e1 = op(e3, e3)) | ~ sP58), inference(cnf_transformation, [], [f128])).
fof(f948, plain, (~ spl66_71 | spl66_8), inference(avatar_split_clause, [], [f313, f619, f942])).
fof(f313, plain, ((e3 = op(e3, e2)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e0 = op(e0, e2)) & (op(e0, e0) = e2) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP59), inference(nnf_transformation, [], [f113])).
fof(f947, plain, (~ spl66_71 | ~ spl66_3), inference(avatar_split_clause, [], [f314, f598, f942])).
fof(f314, plain, (~ (e2 = op(e3, e3)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f946, plain, (~ spl66_71 | spl66_63), inference(avatar_split_clause, [], [f315, f853, f942])).
fof(f315, plain, ((op(e0, e0) = e2) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f945, plain, (~ spl66_71 | ~ spl66_53), inference(avatar_split_clause, [], [f316, f811, f942])).
fof(f316, plain, (~ (e0 = op(e0, e2)) | ~ sP59), inference(cnf_transformation, [], [f127])).
fof(f940, plain, (~ spl66_70 | spl66_8), inference(avatar_split_clause, [], [f309, f619, f934])).
fof(f309, plain, ((e3 = op(e3, e2)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e1 = op(e1, e2)) & (e2 = op(e1, e1)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP60), inference(nnf_transformation, [], [f114])).
fof(f939, plain, (~ spl66_70 | ~ spl66_3), inference(avatar_split_clause, [], [f310, f598, f934])).
fof(f310, plain, (~ (e2 = op(e3, e3)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f938, plain, (~ spl66_70 | spl66_43), inference(avatar_split_clause, [], [f311, f768, f934])).
fof(f311, plain, ((e2 = op(e1, e1)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f937, plain, (~ spl66_70 | ~ spl66_38), inference(avatar_split_clause, [], [f312, f747, f934])).
fof(f312, plain, (~ (e1 = op(e1, e2)) | ~ sP60), inference(cnf_transformation, [], [f126])).
fof(f932, plain, (~ spl66_69 | spl66_8), inference(avatar_split_clause, [], [f305, f619, f926])).
fof(f305, plain, ((e3 = op(e3, e2)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e2 = op(e2, e2)) & (e2 = op(e2, e2)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP61), inference(nnf_transformation, [], [f115])).
fof(f930, plain, (~ spl66_69 | spl66_23), inference(avatar_split_clause, [], [f307, f683, f926])).
fof(f307, plain, ((e2 = op(e2, e2)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f929, plain, (~ spl66_69 | ~ spl66_23), inference(avatar_split_clause, [], [f308, f683, f926])).
fof(f308, plain, (~ (e2 = op(e2, e2)) | ~ sP61), inference(cnf_transformation, [], [f125])).
fof(f924, plain, (~ spl66_68 | spl66_8), inference(avatar_split_clause, [], [f301, f619, f918])).
fof(f301, plain, ((e3 = op(e3, e2)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e3 = op(e3, e2)) & (e2 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & (e3 = op(e3, e2))) | ~ sP62), inference(nnf_transformation, [], [f116])).
fof(f923, plain, (~ spl66_68 | ~ spl66_3), inference(avatar_split_clause, [], [f302, f598, f918])).
fof(f302, plain, (~ (e2 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f922, plain, (~ spl66_68 | spl66_3), inference(avatar_split_clause, [], [f303, f598, f918])).
fof(f303, plain, ((e2 = op(e3, e3)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f921, plain, (~ spl66_68 | ~ spl66_8), inference(avatar_split_clause, [], [f304, f619, f918])).
fof(f304, plain, (~ (e3 = op(e3, e2)) | ~ sP62), inference(cnf_transformation, [], [f124])).
fof(f916, plain, (~ spl66_67 | spl66_4), inference(avatar_split_clause, [], [f297, f602, f910])).
fof(f297, plain, ((e3 = op(e3, e3)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e0 = op(e0, e3)) & (op(e0, e0) = e3) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP63), inference(nnf_transformation, [], [f117])).
fof(f915, plain, (~ spl66_67 | ~ spl66_4), inference(avatar_split_clause, [], [f298, f602, f910])).
fof(f298, plain, (~ (e3 = op(e3, e3)) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f914, plain, (~ spl66_67 | spl66_64), inference(avatar_split_clause, [], [f299, f857, f910])).
fof(f299, plain, ((op(e0, e0) = e3) | ~ sP63), inference(cnf_transformation, [], [f123])).
fof(f908, plain, (~ spl66_66 | spl66_4), inference(avatar_split_clause, [], [f293, f602, f902])).
fof(f293, plain, ((e3 = op(e3, e3)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e1 = op(e1, e3)) & (e3 = op(e1, e1)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP64), inference(nnf_transformation, [], [f118])).
fof(f907, plain, (~ spl66_66 | ~ spl66_4), inference(avatar_split_clause, [], [f294, f602, f902])).
fof(f294, plain, (~ (e3 = op(e3, e3)) | ~ sP64), inference(cnf_transformation, [], [f122])).
fof(f900, plain, (~ spl66_65 | spl66_4), inference(avatar_split_clause, [], [f289, f602, f894])).
fof(f289, plain, ((e3 = op(e3, e3)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e2 = op(e2, e3)) & (e3 = op(e2, e2)) & ~ (e3 = op(e3, e3)) & (e3 = op(e3, e3))) | ~ sP65), inference(nnf_transformation, [], [f119])).
fof(f899, plain, (~ spl66_65 | ~ spl66_4), inference(avatar_split_clause, [], [f290, f602, f894])).
fof(f290, plain, (~ (e3 = op(e3, e3)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f898, plain, (~ spl66_65 | spl66_24), inference(avatar_split_clause, [], [f291, f687, f894])).
fof(f291, plain, ((e3 = op(e2, e2)) | ~ sP65), inference(cnf_transformation, [], [f121])).
fof(f892, plain, (spl66_61 | spl66_57 | spl66_53 | spl66_49), inference(avatar_split_clause, [], [f203, f794, f811, f828, f845])).
fof(f203, plain, ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f2, plain, (((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)) & ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)) & ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)) & ((e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (op(e0, e0) = e2)) & ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)) & ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)) & ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax2)).
fof(f891, plain, (spl66_61 | spl66_45 | spl66_29 | spl66_13), inference(avatar_split_clause, [], [f204, f641, f709, f777, f845])).
fof(f204, plain, ((e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f2])).
fof(f890, plain, (spl66_62 | spl66_58 | spl66_54 | spl66_50), inference(avatar_split_clause, [], [f205, f798, f815, f832, f849])).
fof(f205, plain, ((e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f889, plain, (spl66_62 | spl66_46 | spl66_30 | spl66_14), inference(avatar_split_clause, [], [f206, f645, f713, f781, f849])).
fof(f206, plain, ((e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (op(e0, e0) = e1)), inference(cnf_transformation, [], [f2])).
fof(f887, plain, (spl66_63 | spl66_47 | spl66_31 | spl66_15), inference(avatar_split_clause, [], [f208, f649, f717, f785, f853])).
fof(f208, plain, ((e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (op(e0, e0) = e2)), inference(cnf_transformation, [], [f2])).
fof(f886, plain, (spl66_64 | spl66_60 | spl66_56 | spl66_52), inference(avatar_split_clause, [], [f209, f806, f823, f840, f857])).
fof(f209, plain, ((e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f885, plain, (spl66_64 | spl66_48 | spl66_32 | spl66_16), inference(avatar_split_clause, [], [f210, f653, f721, f789, f857])).
fof(f210, plain, ((e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (op(e0, e0) = e3)), inference(cnf_transformation, [], [f2])).
fof(f883, plain, (spl66_57 | spl66_41 | spl66_25 | spl66_9), inference(avatar_split_clause, [], [f212, f624, f692, f760, f828])).
fof(f212, plain, ((e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f882, plain, (spl66_46 | spl66_42 | spl66_38 | spl66_34), inference(avatar_split_clause, [], [f213, f730, f747, f764, f781])).
fof(f213, plain, ((e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f880, plain, (spl66_47 | spl66_43 | spl66_39 | spl66_35), inference(avatar_split_clause, [], [f215, f734, f751, f768, f785])).
fof(f215, plain, ((e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f879, plain, (spl66_59 | spl66_43 | spl66_27 | spl66_11), inference(avatar_split_clause, [], [f216, f632, f700, f768, f836])).
fof(f216, plain, ((e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f878, plain, (spl66_48 | spl66_44 | spl66_40 | spl66_36), inference(avatar_split_clause, [], [f217, f738, f755, f772, f789])).
fof(f217, plain, ((e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))), inference(cnf_transformation, [], [f2])).
fof(f877, plain, (spl66_60 | spl66_44 | spl66_28 | spl66_12), inference(avatar_split_clause, [], [f218, f636, f704, f772, f840])).
fof(f218, plain, ((e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))), inference(cnf_transformation, [], [f2])).
fof(f876, plain, (spl66_29 | spl66_25 | spl66_21 | spl66_17), inference(avatar_split_clause, [], [f219, f658, f675, f692, f709])).
fof(f219, plain, ((e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f875, plain, (spl66_53 | spl66_37 | spl66_21 | spl66_5), inference(avatar_split_clause, [], [f220, f607, f675, f743, f811])).
fof(f220, plain, ((e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f874, plain, (spl66_30 | spl66_26 | spl66_22 | spl66_18), inference(avatar_split_clause, [], [f221, f662, f679, f696, f713])).
fof(f221, plain, ((e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f873, plain, (spl66_54 | spl66_38 | spl66_22 | spl66_6), inference(avatar_split_clause, [], [f222, f611, f679, f747, f815])).
fof(f222, plain, ((e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f872, plain, (spl66_31 | spl66_27 | spl66_23 | spl66_19), inference(avatar_split_clause, [], [f223, f666, f683, f700, f717])).
fof(f223, plain, ((e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f870, plain, (spl66_32 | spl66_28 | spl66_24 | spl66_20), inference(avatar_split_clause, [], [f225, f670, f687, f704, f721])).
fof(f225, plain, ((e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))), inference(cnf_transformation, [], [f2])).
fof(f869, plain, (spl66_56 | spl66_40 | spl66_24 | spl66_8), inference(avatar_split_clause, [], [f226, f619, f687, f755, f823])).
fof(f226, plain, ((e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))), inference(cnf_transformation, [], [f2])).
fof(f868, plain, (spl66_13 | spl66_9 | spl66_5 | spl66_1), inference(avatar_split_clause, [], [f227, f590, f607, f624, f641])).
fof(f227, plain, ((e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f867, plain, (spl66_49 | spl66_33 | spl66_17 | spl66_1), inference(avatar_split_clause, [], [f228, f590, f658, f726, f794])).
fof(f228, plain, ((e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f866, plain, (spl66_14 | spl66_10 | spl66_6 | spl66_2), inference(avatar_split_clause, [], [f229, f594, f611, f628, f645])).
fof(f229, plain, ((e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f865, plain, (spl66_50 | spl66_34 | spl66_18 | spl66_2), inference(avatar_split_clause, [], [f230, f594, f662, f730, f798])).
fof(f230, plain, ((e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f864, plain, (spl66_15 | spl66_11 | spl66_7 | spl66_3), inference(avatar_split_clause, [], [f231, f598, f615, f632, f649])).
fof(f231, plain, ((e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f863, plain, (spl66_51 | spl66_35 | spl66_19 | spl66_3), inference(avatar_split_clause, [], [f232, f598, f666, f734, f802])).
fof(f232, plain, ((e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f862, plain, (spl66_16 | spl66_12 | spl66_8 | spl66_4), inference(avatar_split_clause, [], [f233, f602, f619, f636, f653])).
fof(f233, plain, ((e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))), inference(cnf_transformation, [], [f2])).
fof(f861, plain, (spl66_52 | spl66_36 | spl66_20 | spl66_4), inference(avatar_split_clause, [], [f234, f602, f670, f738, f806])).
fof(f234, plain, ((e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))), inference(cnf_transformation, [], [f2])).
fof(f860, plain, (spl66_61 | spl66_62 | spl66_63 | spl66_64), inference(avatar_split_clause, [], [f187, f857, f853, f849, f845])).
fof(f187, plain, ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0))), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((op(e0, e0) = e3) | (op(e0, e0) = e2) | (op(e0, e0) = e1) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG156+1.p', ax1)).
fof(f843, plain, (spl66_57 | spl66_58 | spl66_59 | spl66_60), inference(avatar_split_clause, [], [f188, f840, f836, f832, f828])).
fof(f188, plain, ((e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))), inference(cnf_transformation, [], [f1])).
fof(f826, plain, (spl66_53 | spl66_54 | spl66_55 | spl66_56), inference(avatar_split_clause, [], [f189, f823, f819, f815, f811])).
fof(f189, plain, ((e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))), inference(cnf_transformation, [], [f1])).
fof(f809, plain, (spl66_49 | spl66_50 | spl66_51 | spl66_52), inference(avatar_split_clause, [], [f190, f806, f802, f798, f794])).
fof(f190, plain, ((e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))), inference(cnf_transformation, [], [f1])).
fof(f792, plain, (spl66_45 | spl66_46 | spl66_47 | spl66_48), inference(avatar_split_clause, [], [f191, f789, f785, f781, f777])).
fof(f191, plain, ((e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))), inference(cnf_transformation, [], [f1])).
fof(f775, plain, (spl66_41 | spl66_42 | spl66_43 | spl66_44), inference(avatar_split_clause, [], [f192, f772, f768, f764, f760])).
fof(f192, plain, ((e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))), inference(cnf_transformation, [], [f1])).
fof(f758, plain, (spl66_37 | spl66_38 | spl66_39 | spl66_40), inference(avatar_split_clause, [], [f193, f755, f751, f747, f743])).
fof(f193, plain, ((e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))), inference(cnf_transformation, [], [f1])).
fof(f741, plain, (spl66_33 | spl66_34 | spl66_35 | spl66_36), inference(avatar_split_clause, [], [f194, f738, f734, f730, f726])).
fof(f194, plain, ((e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))), inference(cnf_transformation, [], [f1])).
fof(f724, plain, (spl66_29 | spl66_30 | spl66_31 | spl66_32), inference(avatar_split_clause, [], [f195, f721, f717, f713, f709])).
fof(f195, plain, ((e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))), inference(cnf_transformation, [], [f1])).
fof(f707, plain, (spl66_25 | spl66_26 | spl66_27 | spl66_28), inference(avatar_split_clause, [], [f196, f704, f700, f696, f692])).
fof(f196, plain, ((e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))), inference(cnf_transformation, [], [f1])).
fof(f690, plain, (spl66_21 | spl66_22 | spl66_23 | spl66_24), inference(avatar_split_clause, [], [f197, f687, f683, f679, f675])).
fof(f197, plain, ((e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))), inference(cnf_transformation, [], [f1])).
fof(f673, plain, (spl66_17 | spl66_18 | spl66_19 | spl66_20), inference(avatar_split_clause, [], [f198, f670, f666, f662, f658])).
fof(f198, plain, ((e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))), inference(cnf_transformation, [], [f1])).
fof(f656, plain, (spl66_13 | spl66_14 | spl66_15 | spl66_16), inference(avatar_split_clause, [], [f199, f653, f649, f645, f641])).
fof(f199, plain, ((e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))), inference(cnf_transformation, [], [f1])).
fof(f639, plain, (spl66_9 | spl66_10 | spl66_11 | spl66_12), inference(avatar_split_clause, [], [f200, f636, f632, f628, f624])).
fof(f200, plain, ((e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))), inference(cnf_transformation, [], [f1])).
fof(f622, plain, (spl66_5 | spl66_6 | spl66_7 | spl66_8), inference(avatar_split_clause, [], [f201, f619, f615, f611, f607])).
fof(f201, plain, ((e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))), inference(cnf_transformation, [], [f1])).
fof(f605, plain, (spl66_1 | spl66_2 | spl66_3 | spl66_4), inference(avatar_split_clause, [], [f202, f602, f598, f594, f590])).
fof(f202, plain, ((e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))), inference(cnf_transformation, [], [f1])).