fof(f2285, plain, $false, inference(subsumption_resolution, [], [f2284, f320])).
fof(f320, plain, (e0 = op(e0, e0)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((e6 = op(e6, e6)) & (e2 = op(e6, e5)) & (e3 = op(e6, e4)) & (e1 = op(e6, e3)) & (e0 = op(e6, e2)) & (e4 = op(e6, e1)) & (e5 = op(e6, e0)) & (e2 = op(e5, e6)) & (e5 = op(e5, e5)) & (e1 = op(e5, e4)) & (e6 = op(e5, e3)) & (e3 = op(e5, e2)) & (e0 = op(e5, e1)) & (e4 = op(e5, e0)) & (e3 = op(e4, e6)) & (e1 = op(e4, e5)) & (e4 = op(e4, e4)) & (e0 = op(e4, e3)) & (e5 = op(e4, e2)) & (e2 = op(e4, e1)) & (e6 = op(e4, e0)) & (e1 = op(e3, e6)) & (e6 = op(e3, e5)) & (e0 = op(e3, e4)) & (e3 = op(e3, e3)) & (e4 = op(e3, e2)) & (e5 = op(e3, e1)) & (e2 = op(e3, e0)) & (e0 = op(e2, e6)) & (e3 = op(e2, e5)) & (e5 = op(e2, e4)) & (e4 = op(e2, e3)) & (e2 = op(e2, e2)) & (e6 = op(e2, e1)) & (e1 = op(e2, e0)) & (e4 = op(e1, e6)) & (e0 = op(e1, e5)) & (e2 = op(e1, e4)) & (e5 = op(e1, e3)) & (e6 = op(e1, e2)) & (e1 = op(e1, e1)) & (e3 = op(e1, e0)) & (e5 = op(e0, e6)) & (e4 = op(e0, e5)) & (e6 = op(e0, e4)) & (e2 = op(e0, e3)) & (e1 = op(e0, e2)) & (e3 = op(e0, e1)) & (e0 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG208+1.p', ax2)).
fof(f2284, plain, ~ (e0 = op(e0, e0)), inference(forward_demodulation, [], [f2283, f320])).
fof(f2283, plain, ~ (e0 = op(op(e0, e0), e0)), inference(forward_demodulation, [], [f2282, f320])).
fof(f2282, plain, ~ (e0 = op(op(op(e0, e0), e0), e0)), inference(subsumption_resolution, [], [f2281, f356])).
fof(f356, plain, (e0 = op(e5, e1)), inference(cnf_transformation, [], [f2])).
fof(f2281, plain, (~ (e0 = op(e5, e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2280, f342])).
fof(f342, plain, (e5 = op(e3, e1)), inference(cnf_transformation, [], [f2])).
fof(f2280, plain, (~ (e0 = op(op(e3, e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2279, f327])).
fof(f327, plain, (e3 = op(e1, e0)), inference(cnf_transformation, [], [f2])).
fof(f2279, plain, (~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2278, f364])).
fof(f364, plain, (e0 = op(e6, e2)), inference(cnf_transformation, [], [f2])).
fof(f2278, plain, (~ (e0 = op(e6, e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2277, f329])).
fof(f329, plain, (e6 = op(e1, e2)), inference(cnf_transformation, [], [f2])).
fof(f2277, plain, (~ (e0 = op(op(e1, e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2276, f334])).
fof(f334, plain, (e1 = op(e2, e0)), inference(cnf_transformation, [], [f2])).
fof(f2276, plain, (~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2275, f351])).
fof(f351, plain, (e0 = op(e4, e3)), inference(cnf_transformation, [], [f2])).
fof(f2275, plain, (~ (e0 = op(e4, e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2274, f337])).
fof(f337, plain, (e4 = op(e2, e3)), inference(cnf_transformation, [], [f2])).
fof(f2274, plain, (~ (e0 = op(op(e2, e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2273, f341])).
fof(f341, plain, (e2 = op(e3, e0)), inference(cnf_transformation, [], [f2])).
fof(f2273, plain, (~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2272, f345])).
fof(f345, plain, (e0 = op(e3, e4)), inference(cnf_transformation, [], [f2])).
fof(f2272, plain, (~ (e0 = op(e3, e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2271, f366])).
fof(f366, plain, (e3 = op(e6, e4)), inference(cnf_transformation, [], [f2])).
fof(f2271, plain, (~ (e0 = op(op(e6, e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2270, f348])).
fof(f348, plain, (e6 = op(e4, e0)), inference(cnf_transformation, [], [f2])).
fof(f2270, plain, (~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2269, f332])).
fof(f332, plain, (e0 = op(e1, e5)), inference(cnf_transformation, [], [f2])).
fof(f2269, plain, (~ (e0 = op(e1, e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2268, f353])).
fof(f353, plain, (e1 = op(e4, e5)), inference(cnf_transformation, [], [f2])).
fof(f2268, plain, (~ (e0 = op(op(e4, e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2267, f355])).
fof(f355, plain, (e4 = op(e5, e0)), inference(cnf_transformation, [], [f2])).
fof(f2267, plain, (~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2266, f340])).
fof(f340, plain, (e0 = op(e2, e6)), inference(cnf_transformation, [], [f2])).
fof(f2266, plain, (~ (e0 = op(e2, e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2265, f361])).
fof(f361, plain, (e2 = op(e5, e6)), inference(cnf_transformation, [], [f2])).
fof(f2265, plain, (~ (e0 = op(op(e5, e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2264, f362])).
fof(f362, plain, (e5 = op(e6, e0)), inference(cnf_transformation, [], [f2])).
fof(f2264, plain, (~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2263, f334])).
fof(f2263, plain, (~ (e1 = op(e2, e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2262, f341])).
fof(f2262, plain, (~ (e1 = op(op(e3, e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2261, f321])).
fof(f321, plain, (e3 = op(e0, e1)), inference(cnf_transformation, [], [f2])).
fof(f2261, plain, (~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2260, f328])).
fof(f328, plain, (e1 = op(e1, e1)), inference(cnf_transformation, [], [f2])).
fof(f2260, plain, (~ (e1 = op(e1, e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2259, f328])).
fof(f2259, plain, (~ (e1 = op(op(e1, e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2258, f328])).
fof(f2258, plain, (~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2257, f322])).
fof(f322, plain, (e1 = op(e0, e2)), inference(cnf_transformation, [], [f2])).
fof(f2257, plain, (~ (e1 = op(e0, e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2256, f364])).
fof(f2256, plain, (~ (e1 = op(op(e6, e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2255, f335])).
fof(f335, plain, (e6 = op(e2, e1)), inference(cnf_transformation, [], [f2])).
fof(f2255, plain, (~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2254, f365])).
fof(f365, plain, (e1 = op(e6, e3)), inference(cnf_transformation, [], [f2])).
fof(f2254, plain, (~ (e1 = op(e6, e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2253, f358])).
fof(f358, plain, (e6 = op(e5, e3)), inference(cnf_transformation, [], [f2])).
fof(f2253, plain, (~ (e1 = op(op(e5, e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2252, f342])).
fof(f2252, plain, (~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2251, f359])).
fof(f359, plain, (e1 = op(e5, e4)), inference(cnf_transformation, [], [f2])).
fof(f2251, plain, (~ (e1 = op(e5, e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2250, f338])).
fof(f338, plain, (e5 = op(e2, e4)), inference(cnf_transformation, [], [f2])).
fof(f2250, plain, (~ (e1 = op(op(e2, e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2249, f349])).
fof(f349, plain, (e2 = op(e4, e1)), inference(cnf_transformation, [], [f2])).
fof(f2249, plain, (~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2248, f353])).
fof(f2248, plain, (~ (e1 = op(e4, e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2247, f325])).
fof(f325, plain, (e4 = op(e0, e5)), inference(cnf_transformation, [], [f2])).
fof(f2247, plain, (~ (e1 = op(op(e0, e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2246, f356])).
fof(f2246, plain, (~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2245, f347])).
fof(f347, plain, (e1 = op(e3, e6)), inference(cnf_transformation, [], [f2])).
fof(f2245, plain, (~ (e1 = op(e3, e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2244, f354])).
fof(f354, plain, (e3 = op(e4, e6)), inference(cnf_transformation, [], [f2])).
fof(f2244, plain, (~ (e1 = op(op(e4, e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2243, f363])).
fof(f363, plain, (e4 = op(e6, e1)), inference(cnf_transformation, [], [f2])).
fof(f2243, plain, (~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2242, f341])).
fof(f2242, plain, (~ (e2 = op(e3, e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2241, f327])).
fof(f2241, plain, (~ (e2 = op(op(e1, e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2240, f322])).
fof(f2240, plain, (~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2239, f349])).
fof(f2239, plain, (~ (e2 = op(e4, e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2238, f363])).
fof(f2238, plain, (~ (e2 = op(op(e6, e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2237, f329])).
fof(f2237, plain, (~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2236, f336])).
fof(f336, plain, (e2 = op(e2, e2)), inference(cnf_transformation, [], [f2])).
fof(f2236, plain, (~ (e2 = op(e2, e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2235, f336])).
fof(f2235, plain, (~ (e2 = op(op(e2, e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2234, f336])).
fof(f2234, plain, (~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2233, f323])).
fof(f323, plain, (e2 = op(e0, e3)), inference(cnf_transformation, [], [f2])).
fof(f2233, plain, (~ (e2 = op(e0, e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2232, f351])).
fof(f2232, plain, (~ (e2 = op(op(e4, e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2231, f343])).
fof(f343, plain, (e4 = op(e3, e2)), inference(cnf_transformation, [], [f2])).
fof(f2231, plain, (~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2230, f331])).
fof(f331, plain, (e2 = op(e1, e4)), inference(cnf_transformation, [], [f2])).
fof(f2230, plain, (~ (e2 = op(e1, e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2229, f359])).
fof(f2229, plain, (~ (e2 = op(op(e5, e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2228, f350])).
fof(f350, plain, (e5 = op(e4, e2)), inference(cnf_transformation, [], [f2])).
fof(f2228, plain, (~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2227, f367])).
fof(f367, plain, (e2 = op(e6, e5)), inference(cnf_transformation, [], [f2])).
fof(f2227, plain, (~ (e2 = op(e6, e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2226, f346])).
fof(f346, plain, (e6 = op(e3, e5)), inference(cnf_transformation, [], [f2])).
fof(f2226, plain, (~ (e2 = op(op(e3, e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2225, f357])).
fof(f357, plain, (e3 = op(e5, e2)), inference(cnf_transformation, [], [f2])).
fof(f2225, plain, (~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2224, f361])).
fof(f2224, plain, (~ (e2 = op(e5, e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2223, f326])).
fof(f326, plain, (e5 = op(e0, e6)), inference(cnf_transformation, [], [f2])).
fof(f2223, plain, (~ (e2 = op(op(e0, e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2222, f364])).
fof(f2222, plain, (~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2221, f327])).
fof(f2221, plain, (~ (e3 = op(e1, e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2220, f334])).
fof(f2220, plain, (~ (e3 = op(op(e2, e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2219, f323])).
fof(f2219, plain, (~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2218, f321])).
fof(f2218, plain, (~ (e3 = op(e0, e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2217, f356])).
fof(f2217, plain, (~ (e3 = op(op(e5, e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2216, f330])).
fof(f330, plain, (e5 = op(e1, e3)), inference(cnf_transformation, [], [f2])).
fof(f2216, plain, (~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2215, f357])).
fof(f2215, plain, (~ (e3 = op(e5, e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2214, f350])).
fof(f2214, plain, (~ (e3 = op(op(e4, e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2213, f337])).
fof(f2213, plain, (~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2212, f344])).
fof(f344, plain, (e3 = op(e3, e3)), inference(cnf_transformation, [], [f2])).
fof(f2212, plain, (~ (e3 = op(e3, e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2211, f344])).
fof(f2211, plain, (~ (e3 = op(op(e3, e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2210, f344])).
fof(f2210, plain, (~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2209, f366])).
fof(f2209, plain, (~ (e3 = op(e6, e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2208, f324])).
fof(f324, plain, (e6 = op(e0, e4)), inference(cnf_transformation, [], [f2])).
fof(f2208, plain, (~ (e3 = op(op(e0, e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2207, f351])).
fof(f2207, plain, (~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2206, f339])).
fof(f339, plain, (e3 = op(e2, e5)), inference(cnf_transformation, [], [f2])).
fof(f2206, plain, (~ (e3 = op(e2, e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2205, f367])).
fof(f2205, plain, (~ (e3 = op(op(e6, e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2204, f358])).
fof(f2204, plain, (~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2203, f354])).
fof(f2203, plain, (~ (e3 = op(e4, e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2202, f333])).
fof(f333, plain, (e4 = op(e1, e6)), inference(cnf_transformation, [], [f2])).
fof(f2202, plain, (~ (e3 = op(op(e1, e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2201, f365])).
fof(f2201, plain, (~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2200, f355])).
fof(f2200, plain, (~ (e4 = op(e5, e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2199, f362])).
fof(f2199, plain, (~ (e4 = op(op(e6, e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2198, f324])).
fof(f2198, plain, (~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2197, f363])).
fof(f2197, plain, (~ (e4 = op(e6, e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2196, f335])).
fof(f2196, plain, (~ (e4 = op(op(e2, e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2195, f331])).
fof(f2195, plain, (~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2194, f343])).
fof(f2194, plain, (~ (e4 = op(e3, e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2193, f357])).
fof(f2193, plain, (~ (e4 = op(op(e5, e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2192, f338])).
fof(f2192, plain, (~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2191, f337])).
fof(f2191, plain, (~ (e4 = op(e2, e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2190, f323])).
fof(f2190, plain, (~ (e4 = op(op(e0, e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2189, f345])).
fof(f2189, plain, (~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2188, f352])).
fof(f352, plain, (e4 = op(e4, e4)), inference(cnf_transformation, [], [f2])).
fof(f2188, plain, (~ (e4 = op(e4, e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2187, f352])).
fof(f2187, plain, (~ (e4 = op(op(e4, e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2186, f352])).
fof(f2186, plain, (~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2185, f325])).
fof(f2185, plain, (~ (e4 = op(e0, e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2184, f332])).
fof(f2184, plain, (~ (e4 = op(op(e1, e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2183, f359])).
fof(f2183, plain, (~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2182, f333])).
fof(f2182, plain, (~ (e4 = op(e1, e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2181, f347])).
fof(f2181, plain, (~ (e4 = op(op(e3, e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2180, f366])).
fof(f2180, plain, (~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2179, f362])).
fof(f2179, plain, (~ (e5 = op(e6, e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2178, f348])).
fof(f2178, plain, (~ (e5 = op(op(e4, e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2177, f325])).
fof(f2177, plain, (~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2176, f342])).
fof(f2176, plain, (~ (e5 = op(e3, e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2175, f321])).
fof(f2175, plain, (~ (e5 = op(op(e0, e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2174, f332])).
fof(f2174, plain, (~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2173, f350])).
fof(f2173, plain, (~ (e5 = op(e4, e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2172, f343])).
fof(f2172, plain, (~ (e5 = op(op(e3, e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2171, f339])).
fof(f2171, plain, (~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2170, f330])).
fof(f2170, plain, (~ (e5 = op(e1, e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2169, f365])).
fof(f2169, plain, (~ (e5 = op(op(e6, e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2168, f346])).
fof(f2168, plain, (~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2167, f338])).
fof(f2167, plain, (~ (e5 = op(e2, e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2166, f331])).
fof(f2166, plain, (~ (e5 = op(op(e1, e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2165, f353])).
fof(f2165, plain, (~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2164, f360])).
fof(f360, plain, (e5 = op(e5, e5)), inference(cnf_transformation, [], [f2])).
fof(f2164, plain, (~ (e5 = op(e5, e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2163, f360])).
fof(f2163, plain, (~ (e5 = op(op(e5, e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2162, f360])).
fof(f2162, plain, (~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2161, f326])).
fof(f2161, plain, (~ (e5 = op(e0, e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2160, f340])).
fof(f2160, plain, (~ (e5 = op(op(e2, e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2159, f367])).
fof(f2159, plain, (~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2158, f348])).
fof(f2158, plain, (~ (e6 = op(e4, e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2157, f355])).
fof(f2157, plain, (~ (e6 = op(op(e5, e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2156, f326])).
fof(f2156, plain, (~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2155, f335])).
fof(f2155, plain, (~ (e6 = op(e2, e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2154, f349])).
fof(f2154, plain, (~ (e6 = op(op(e4, e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2153, f333])).
fof(f2153, plain, (~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2152, f329])).
fof(f2152, plain, (~ (e6 = op(e1, e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2151, f322])).
fof(f2151, plain, (~ (e6 = op(op(e0, e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2150, f340])).
fof(f2150, plain, (~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2149, f358])).
fof(f2149, plain, (~ (e6 = op(e5, e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2148, f330])).
fof(f2148, plain, (~ (e6 = op(op(e1, e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2147, f347])).
fof(f2147, plain, (~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2146, f324])).
fof(f2146, plain, (~ (e6 = op(e0, e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2145, f345])).
fof(f2145, plain, (~ (e6 = op(op(e3, e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2144, f354])).
fof(f2144, plain, (~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2143, f346])).
fof(f2143, plain, (~ (e6 = op(e3, e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2142, f339])).
fof(f2142, plain, (~ (e6 = op(op(e2, e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2141, f361])).
fof(f2141, plain, (~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2140, f368])).
fof(f368, plain, (e6 = op(e6, e6)), inference(cnf_transformation, [], [f2])).
fof(f2140, plain, (~ (e6 = op(e6, e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2139, f368])).
fof(f2139, plain, (~ (e6 = op(op(e6, e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f2138, f368])).
fof(f2138, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f2137, f320])).
fof(f2137, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2136, f328])).
fof(f2136, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2135, f336])).
fof(f2135, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2134, f344])).
fof(f2134, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2133, f352])).
fof(f2133, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2132, f360])).
fof(f2132, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2131, f1985])).
fof(f1985, plain, ~ sP0, inference(subsumption_resolution, [], [f1384, f320])).
fof(f1384, plain, (~ (e0 = op(e0, e0)) | ~ sP0), inference(cnf_transformation, [], [f298])).
fof(f298, plain, ((~ (e6 = op(e0, e0)) & ~ (e5 = op(e0, e0)) & ~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f6])).
fof(f6, plain, ((~ (e6 = op(e0, e0)) & ~ (e5 = op(e0, e0)) & ~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e6])).
fof(e6, plain, (sP0 <=> (~ (e6 = op(e0, e0)) & ~ (e5 = op(e0, e0)) & ~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f2131, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2130, f1978])).
fof(f1978, plain, ~ sP1, inference(subsumption_resolution, [], [f1380, f321])).
fof(f1380, plain, (~ (e3 = op(e0, e1)) | ~ sP1), inference(cnf_transformation, [], [f297])).
fof(f297, plain, ((~ (e6 = op(e0, e1)) & ~ (e5 = op(e0, e1)) & ~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP1), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ((~ (e6 = op(e0, e1)) & ~ (e5 = op(e0, e1)) & ~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP1), inference(usedef, [], [e7])).
fof(e7, plain, (sP1 <=> (~ (e6 = op(e0, e1)) & ~ (e5 = op(e0, e1)) & ~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f2130, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2129, f1974])).
fof(f1974, plain, ~ sP2, inference(subsumption_resolution, [], [f1371, f322])).
fof(f1371, plain, (~ (e1 = op(e0, e2)) | ~ sP2), inference(cnf_transformation, [], [f296])).
fof(f296, plain, ((~ (e6 = op(e0, e2)) & ~ (e5 = op(e0, e2)) & ~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP2), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ((~ (e6 = op(e0, e2)) & ~ (e5 = op(e0, e2)) & ~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP2), inference(usedef, [], [e8])).
fof(e8, plain, (sP2 <=> (~ (e6 = op(e0, e2)) & ~ (e5 = op(e0, e2)) & ~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f2129, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2128, f1968])).
fof(f1968, plain, ~ sP3, inference(subsumption_resolution, [], [f1365, f323])).
fof(f1365, plain, (~ (e2 = op(e0, e3)) | ~ sP3), inference(cnf_transformation, [], [f295])).
fof(f295, plain, ((~ (e6 = op(e0, e3)) & ~ (e5 = op(e0, e3)) & ~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP3), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ((~ (e6 = op(e0, e3)) & ~ (e5 = op(e0, e3)) & ~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP3), inference(usedef, [], [e9])).
fof(e9, plain, (sP3 <=> (~ (e6 = op(e0, e3)) & ~ (e5 = op(e0, e3)) & ~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f2128, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2127, f1963])).
fof(f1963, plain, ~ sP4, inference(subsumption_resolution, [], [f1362, f324])).
fof(f1362, plain, (~ (e6 = op(e0, e4)) | ~ sP4), inference(cnf_transformation, [], [f294])).
fof(f294, plain, ((~ (e6 = op(e0, e4)) & ~ (e5 = op(e0, e4)) & ~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP4), inference(nnf_transformation, [], [f10])).
fof(f10, plain, ((~ (e6 = op(e0, e4)) & ~ (e5 = op(e0, e4)) & ~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP4), inference(usedef, [], [e10])).
fof(e10, plain, (sP4 <=> (~ (e6 = op(e0, e4)) & ~ (e5 = op(e0, e4)) & ~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f2127, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2126, f1962])).
fof(f1962, plain, ~ sP5, inference(subsumption_resolution, [], [f1353, f325])).
fof(f1353, plain, (~ (e4 = op(e0, e5)) | ~ sP5), inference(cnf_transformation, [], [f293])).
fof(f293, plain, ((~ (e6 = op(e0, e5)) & ~ (e5 = op(e0, e5)) & ~ (e4 = op(e0, e5)) & ~ (e3 = op(e0, e5)) & ~ (e2 = op(e0, e5)) & ~ (e1 = op(e0, e5)) & ~ (e0 = op(e0, e5))) | ~ sP5), inference(nnf_transformation, [], [f11])).
fof(f11, plain, ((~ (e6 = op(e0, e5)) & ~ (e5 = op(e0, e5)) & ~ (e4 = op(e0, e5)) & ~ (e3 = op(e0, e5)) & ~ (e2 = op(e0, e5)) & ~ (e1 = op(e0, e5)) & ~ (e0 = op(e0, e5))) | ~ sP5), inference(usedef, [], [e11])).
fof(e11, plain, (sP5 <=> (~ (e6 = op(e0, e5)) & ~ (e5 = op(e0, e5)) & ~ (e4 = op(e0, e5)) & ~ (e3 = op(e0, e5)) & ~ (e2 = op(e0, e5)) & ~ (e1 = op(e0, e5)) & ~ (e0 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f2126, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2125, f1959])).
fof(f1959, plain, ~ sP6, inference(subsumption_resolution, [], [f1347, f326])).
fof(f1347, plain, (~ (e5 = op(e0, e6)) | ~ sP6), inference(cnf_transformation, [], [f292])).
fof(f292, plain, ((~ (e6 = op(e0, e6)) & ~ (e5 = op(e0, e6)) & ~ (e4 = op(e0, e6)) & ~ (e3 = op(e0, e6)) & ~ (e2 = op(e0, e6)) & ~ (e1 = op(e0, e6)) & ~ (e0 = op(e0, e6))) | ~ sP6), inference(nnf_transformation, [], [f12])).
fof(f12, plain, ((~ (e6 = op(e0, e6)) & ~ (e5 = op(e0, e6)) & ~ (e4 = op(e0, e6)) & ~ (e3 = op(e0, e6)) & ~ (e2 = op(e0, e6)) & ~ (e1 = op(e0, e6)) & ~ (e0 = op(e0, e6))) | ~ sP6), inference(usedef, [], [e12])).
fof(e12, plain, (sP6 <=> (~ (e6 = op(e0, e6)) & ~ (e5 = op(e0, e6)) & ~ (e4 = op(e0, e6)) & ~ (e3 = op(e0, e6)) & ~ (e2 = op(e0, e6)) & ~ (e1 = op(e0, e6)) & ~ (e0 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f2125, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2124, f1957])).
fof(f1957, plain, ~ sP7, inference(subsumption_resolution, [], [f1338, f327])).
fof(f1338, plain, (~ (e3 = op(e1, e0)) | ~ sP7), inference(cnf_transformation, [], [f291])).
fof(f291, plain, ((~ (e6 = op(e1, e0)) & ~ (e5 = op(e1, e0)) & ~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP7), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ((~ (e6 = op(e1, e0)) & ~ (e5 = op(e1, e0)) & ~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP7), inference(usedef, [], [e13])).
fof(e13, plain, (sP7 <=> (~ (e6 = op(e1, e0)) & ~ (e5 = op(e1, e0)) & ~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f2124, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2123, f1953])).
fof(f1953, plain, ~ sP8, inference(subsumption_resolution, [], [f1329, f328])).
fof(f1329, plain, (~ (e1 = op(e1, e1)) | ~ sP8), inference(cnf_transformation, [], [f290])).
fof(f290, plain, ((~ (e6 = op(e1, e1)) & ~ (e5 = op(e1, e1)) & ~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP8), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ((~ (e6 = op(e1, e1)) & ~ (e5 = op(e1, e1)) & ~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP8), inference(usedef, [], [e14])).
fof(e14, plain, (sP8 <=> (~ (e6 = op(e1, e1)) & ~ (e5 = op(e1, e1)) & ~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f2123, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2122, f1947])).
fof(f1947, plain, ~ sP9, inference(subsumption_resolution, [], [f1327, f329])).
fof(f1327, plain, (~ (e6 = op(e1, e2)) | ~ sP9), inference(cnf_transformation, [], [f289])).
fof(f289, plain, ((~ (e6 = op(e1, e2)) & ~ (e5 = op(e1, e2)) & ~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP9), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ((~ (e6 = op(e1, e2)) & ~ (e5 = op(e1, e2)) & ~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP9), inference(usedef, [], [e15])).
fof(e15, plain, (sP9 <=> (~ (e6 = op(e1, e2)) & ~ (e5 = op(e1, e2)) & ~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f2122, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2121, f1946])).
fof(f1946, plain, ~ sP10, inference(subsumption_resolution, [], [f1319, f330])).
fof(f1319, plain, (~ (e5 = op(e1, e3)) | ~ sP10), inference(cnf_transformation, [], [f288])).
fof(f288, plain, ((~ (e6 = op(e1, e3)) & ~ (e5 = op(e1, e3)) & ~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP10), inference(nnf_transformation, [], [f16])).
fof(f16, plain, ((~ (e6 = op(e1, e3)) & ~ (e5 = op(e1, e3)) & ~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP10), inference(usedef, [], [e16])).
fof(e16, plain, (sP10 <=> (~ (e6 = op(e1, e3)) & ~ (e5 = op(e1, e3)) & ~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f2121, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2120, f1944])).
fof(f1944, plain, ~ sP11, inference(subsumption_resolution, [], [f1309, f331])).
fof(f1309, plain, (~ (e2 = op(e1, e4)) | ~ sP11), inference(cnf_transformation, [], [f287])).
fof(f287, plain, ((~ (e6 = op(e1, e4)) & ~ (e5 = op(e1, e4)) & ~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP11), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ((~ (e6 = op(e1, e4)) & ~ (e5 = op(e1, e4)) & ~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP11), inference(usedef, [], [e17])).
fof(e17, plain, (sP11 <=> (~ (e6 = op(e1, e4)) & ~ (e5 = op(e1, e4)) & ~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f2120, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2119, f1939])).
fof(f1939, plain, ~ sP12, inference(subsumption_resolution, [], [f1300, f332])).
fof(f1300, plain, (~ (e0 = op(e1, e5)) | ~ sP12), inference(cnf_transformation, [], [f286])).
fof(f286, plain, ((~ (e6 = op(e1, e5)) & ~ (e5 = op(e1, e5)) & ~ (e4 = op(e1, e5)) & ~ (e3 = op(e1, e5)) & ~ (e2 = op(e1, e5)) & ~ (e1 = op(e1, e5)) & ~ (e0 = op(e1, e5))) | ~ sP12), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ((~ (e6 = op(e1, e5)) & ~ (e5 = op(e1, e5)) & ~ (e4 = op(e1, e5)) & ~ (e3 = op(e1, e5)) & ~ (e2 = op(e1, e5)) & ~ (e1 = op(e1, e5)) & ~ (e0 = op(e1, e5))) | ~ sP12), inference(usedef, [], [e18])).
fof(e18, plain, (sP12 <=> (~ (e6 = op(e1, e5)) & ~ (e5 = op(e1, e5)) & ~ (e4 = op(e1, e5)) & ~ (e3 = op(e1, e5)) & ~ (e2 = op(e1, e5)) & ~ (e1 = op(e1, e5)) & ~ (e0 = op(e1, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f2119, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2118, f1932])).
fof(f1932, plain, ~ sP13, inference(subsumption_resolution, [], [f1297, f333])).
fof(f1297, plain, (~ (e4 = op(e1, e6)) | ~ sP13), inference(cnf_transformation, [], [f285])).
fof(f285, plain, ((~ (e6 = op(e1, e6)) & ~ (e5 = op(e1, e6)) & ~ (e4 = op(e1, e6)) & ~ (e3 = op(e1, e6)) & ~ (e2 = op(e1, e6)) & ~ (e1 = op(e1, e6)) & ~ (e0 = op(e1, e6))) | ~ sP13), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ((~ (e6 = op(e1, e6)) & ~ (e5 = op(e1, e6)) & ~ (e4 = op(e1, e6)) & ~ (e3 = op(e1, e6)) & ~ (e2 = op(e1, e6)) & ~ (e1 = op(e1, e6)) & ~ (e0 = op(e1, e6))) | ~ sP13), inference(usedef, [], [e19])).
fof(e19, plain, (sP13 <=> (~ (e6 = op(e1, e6)) & ~ (e5 = op(e1, e6)) & ~ (e4 = op(e1, e6)) & ~ (e3 = op(e1, e6)) & ~ (e2 = op(e1, e6)) & ~ (e1 = op(e1, e6)) & ~ (e0 = op(e1, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f2118, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2117, f1929])).
fof(f1929, plain, ~ sP14, inference(subsumption_resolution, [], [f1287, f334])).
fof(f1287, plain, (~ (e1 = op(e2, e0)) | ~ sP14), inference(cnf_transformation, [], [f284])).
fof(f284, plain, ((~ (e6 = op(e2, e0)) & ~ (e5 = op(e2, e0)) & ~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP14), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ((~ (e6 = op(e2, e0)) & ~ (e5 = op(e2, e0)) & ~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP14), inference(usedef, [], [e20])).
fof(e20, plain, (sP14 <=> (~ (e6 = op(e2, e0)) & ~ (e5 = op(e2, e0)) & ~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f2117, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2116, f1923])).
fof(f1923, plain, ~ sP15, inference(subsumption_resolution, [], [f1285, f335])).
fof(f1285, plain, (~ (e6 = op(e2, e1)) | ~ sP15), inference(cnf_transformation, [], [f283])).
fof(f283, plain, ((~ (e6 = op(e2, e1)) & ~ (e5 = op(e2, e1)) & ~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP15), inference(nnf_transformation, [], [f21])).
fof(f21, plain, ((~ (e6 = op(e2, e1)) & ~ (e5 = op(e2, e1)) & ~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP15), inference(usedef, [], [e21])).
fof(e21, plain, (sP15 <=> (~ (e6 = op(e2, e1)) & ~ (e5 = op(e2, e1)) & ~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f2116, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2115, f1922])).
fof(f1922, plain, ~ sP16, inference(subsumption_resolution, [], [f1274, f336])).
fof(f1274, plain, (~ (e2 = op(e2, e2)) | ~ sP16), inference(cnf_transformation, [], [f282])).
fof(f282, plain, ((~ (e6 = op(e2, e2)) & ~ (e5 = op(e2, e2)) & ~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP16), inference(nnf_transformation, [], [f22])).
fof(f22, plain, ((~ (e6 = op(e2, e2)) & ~ (e5 = op(e2, e2)) & ~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP16), inference(usedef, [], [e22])).
fof(e22, plain, (sP16 <=> (~ (e6 = op(e2, e2)) & ~ (e5 = op(e2, e2)) & ~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f2115, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2114, f1917])).
fof(f1917, plain, ~ sP17, inference(subsumption_resolution, [], [f1269, f337])).
fof(f1269, plain, (~ (e4 = op(e2, e3)) | ~ sP17), inference(cnf_transformation, [], [f281])).
fof(f281, plain, ((~ (e6 = op(e2, e3)) & ~ (e5 = op(e2, e3)) & ~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP17), inference(nnf_transformation, [], [f23])).
fof(f23, plain, ((~ (e6 = op(e2, e3)) & ~ (e5 = op(e2, e3)) & ~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP17), inference(usedef, [], [e23])).
fof(e23, plain, (sP17 <=> (~ (e6 = op(e2, e3)) & ~ (e5 = op(e2, e3)) & ~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f2114, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2113, f1914])).
fof(f1914, plain, ~ sP18, inference(subsumption_resolution, [], [f1263, f338])).
fof(f1263, plain, (~ (e5 = op(e2, e4)) | ~ sP18), inference(cnf_transformation, [], [f280])).
fof(f280, plain, ((~ (e6 = op(e2, e4)) & ~ (e5 = op(e2, e4)) & ~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP18), inference(nnf_transformation, [], [f24])).
fof(f24, plain, ((~ (e6 = op(e2, e4)) & ~ (e5 = op(e2, e4)) & ~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP18), inference(usedef, [], [e24])).
fof(e24, plain, (sP18 <=> (~ (e6 = op(e2, e4)) & ~ (e5 = op(e2, e4)) & ~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f2113, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2112, f1912])).
fof(f1912, plain, ~ sP19, inference(subsumption_resolution, [], [f1254, f339])).
fof(f1254, plain, (~ (e3 = op(e2, e5)) | ~ sP19), inference(cnf_transformation, [], [f279])).
fof(f279, plain, ((~ (e6 = op(e2, e5)) & ~ (e5 = op(e2, e5)) & ~ (e4 = op(e2, e5)) & ~ (e3 = op(e2, e5)) & ~ (e2 = op(e2, e5)) & ~ (e1 = op(e2, e5)) & ~ (e0 = op(e2, e5))) | ~ sP19), inference(nnf_transformation, [], [f25])).
fof(f25, plain, ((~ (e6 = op(e2, e5)) & ~ (e5 = op(e2, e5)) & ~ (e4 = op(e2, e5)) & ~ (e3 = op(e2, e5)) & ~ (e2 = op(e2, e5)) & ~ (e1 = op(e2, e5)) & ~ (e0 = op(e2, e5))) | ~ sP19), inference(usedef, [], [e25])).
fof(e25, plain, (sP19 <=> (~ (e6 = op(e2, e5)) & ~ (e5 = op(e2, e5)) & ~ (e4 = op(e2, e5)) & ~ (e3 = op(e2, e5)) & ~ (e2 = op(e2, e5)) & ~ (e1 = op(e2, e5)) & ~ (e0 = op(e2, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f2112, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2111, f1908])).
fof(f1908, plain, ~ sP20, inference(subsumption_resolution, [], [f1244, f340])).
fof(f1244, plain, (~ (e0 = op(e2, e6)) | ~ sP20), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ((~ (e6 = op(e2, e6)) & ~ (e5 = op(e2, e6)) & ~ (e4 = op(e2, e6)) & ~ (e3 = op(e2, e6)) & ~ (e2 = op(e2, e6)) & ~ (e1 = op(e2, e6)) & ~ (e0 = op(e2, e6))) | ~ sP20), inference(nnf_transformation, [], [f26])).
fof(f26, plain, ((~ (e6 = op(e2, e6)) & ~ (e5 = op(e2, e6)) & ~ (e4 = op(e2, e6)) & ~ (e3 = op(e2, e6)) & ~ (e2 = op(e2, e6)) & ~ (e1 = op(e2, e6)) & ~ (e0 = op(e2, e6))) | ~ sP20), inference(usedef, [], [e26])).
fof(e26, plain, (sP20 <=> (~ (e6 = op(e2, e6)) & ~ (e5 = op(e2, e6)) & ~ (e4 = op(e2, e6)) & ~ (e3 = op(e2, e6)) & ~ (e2 = op(e2, e6)) & ~ (e1 = op(e2, e6)) & ~ (e0 = op(e2, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f2111, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2110, f1901])).
fof(f1901, plain, ~ sP21, inference(subsumption_resolution, [], [f1239, f341])).
fof(f1239, plain, (~ (e2 = op(e3, e0)) | ~ sP21), inference(cnf_transformation, [], [f277])).
fof(f277, plain, ((~ (e6 = op(e3, e0)) & ~ (e5 = op(e3, e0)) & ~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP21), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ((~ (e6 = op(e3, e0)) & ~ (e5 = op(e3, e0)) & ~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP21), inference(usedef, [], [e27])).
fof(e27, plain, (sP21 <=> (~ (e6 = op(e3, e0)) & ~ (e5 = op(e3, e0)) & ~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f2110, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2109, f1896])).
fof(f1896, plain, ~ sP22, inference(subsumption_resolution, [], [f1235, f342])).
fof(f1235, plain, (~ (e5 = op(e3, e1)) | ~ sP22), inference(cnf_transformation, [], [f276])).
fof(f276, plain, ((~ (e6 = op(e3, e1)) & ~ (e5 = op(e3, e1)) & ~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP22), inference(nnf_transformation, [], [f28])).
fof(f28, plain, ((~ (e6 = op(e3, e1)) & ~ (e5 = op(e3, e1)) & ~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP22), inference(usedef, [], [e28])).
fof(e28, plain, (sP22 <=> (~ (e6 = op(e3, e1)) & ~ (e5 = op(e3, e1)) & ~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f2109, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2108, f1894])).
fof(f1894, plain, ~ sP23, inference(subsumption_resolution, [], [f1227, f343])).
fof(f1227, plain, (~ (e4 = op(e3, e2)) | ~ sP23), inference(cnf_transformation, [], [f275])).
fof(f275, plain, ((~ (e6 = op(e3, e2)) & ~ (e5 = op(e3, e2)) & ~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP23), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ((~ (e6 = op(e3, e2)) & ~ (e5 = op(e3, e2)) & ~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP23), inference(usedef, [], [e29])).
fof(e29, plain, (sP23 <=> (~ (e6 = op(e3, e2)) & ~ (e5 = op(e3, e2)) & ~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f2108, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2107, f1891])).
fof(f1891, plain, ~ sP24, inference(subsumption_resolution, [], [f1219, f344])).
fof(f1219, plain, (~ (e3 = op(e3, e3)) | ~ sP24), inference(cnf_transformation, [], [f274])).
fof(f274, plain, ((~ (e6 = op(e3, e3)) & ~ (e5 = op(e3, e3)) & ~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP24), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ((~ (e6 = op(e3, e3)) & ~ (e5 = op(e3, e3)) & ~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP24), inference(usedef, [], [e30])).
fof(e30, plain, (sP24 <=> (~ (e6 = op(e3, e3)) & ~ (e5 = op(e3, e3)) & ~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f2107, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2106, f1887])).
fof(f1887, plain, ~ sP25, inference(subsumption_resolution, [], [f1209, f345])).
fof(f1209, plain, (~ (e0 = op(e3, e4)) | ~ sP25), inference(cnf_transformation, [], [f273])).
fof(f273, plain, ((~ (e6 = op(e3, e4)) & ~ (e5 = op(e3, e4)) & ~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP25), inference(nnf_transformation, [], [f31])).
fof(f31, plain, ((~ (e6 = op(e3, e4)) & ~ (e5 = op(e3, e4)) & ~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP25), inference(usedef, [], [e31])).
fof(e31, plain, (sP25 <=> (~ (e6 = op(e3, e4)) & ~ (e5 = op(e3, e4)) & ~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f2106, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2105, f1880])).
fof(f1880, plain, ~ sP26, inference(subsumption_resolution, [], [f1208, f346])).
fof(f1208, plain, (~ (e6 = op(e3, e5)) | ~ sP26), inference(cnf_transformation, [], [f272])).
fof(f272, plain, ((~ (e6 = op(e3, e5)) & ~ (e5 = op(e3, e5)) & ~ (e4 = op(e3, e5)) & ~ (e3 = op(e3, e5)) & ~ (e2 = op(e3, e5)) & ~ (e1 = op(e3, e5)) & ~ (e0 = op(e3, e5))) | ~ sP26), inference(nnf_transformation, [], [f32])).
fof(f32, plain, ((~ (e6 = op(e3, e5)) & ~ (e5 = op(e3, e5)) & ~ (e4 = op(e3, e5)) & ~ (e3 = op(e3, e5)) & ~ (e2 = op(e3, e5)) & ~ (e1 = op(e3, e5)) & ~ (e0 = op(e3, e5))) | ~ sP26), inference(usedef, [], [e32])).
fof(e32, plain, (sP26 <=> (~ (e6 = op(e3, e5)) & ~ (e5 = op(e3, e5)) & ~ (e4 = op(e3, e5)) & ~ (e3 = op(e3, e5)) & ~ (e2 = op(e3, e5)) & ~ (e1 = op(e3, e5)) & ~ (e0 = op(e3, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f2105, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2104, f1879])).
fof(f1879, plain, ~ sP27, inference(subsumption_resolution, [], [f1196, f347])).
fof(f1196, plain, (~ (e1 = op(e3, e6)) | ~ sP27), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ((~ (e6 = op(e3, e6)) & ~ (e5 = op(e3, e6)) & ~ (e4 = op(e3, e6)) & ~ (e3 = op(e3, e6)) & ~ (e2 = op(e3, e6)) & ~ (e1 = op(e3, e6)) & ~ (e0 = op(e3, e6))) | ~ sP27), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ((~ (e6 = op(e3, e6)) & ~ (e5 = op(e3, e6)) & ~ (e4 = op(e3, e6)) & ~ (e3 = op(e3, e6)) & ~ (e2 = op(e3, e6)) & ~ (e1 = op(e3, e6)) & ~ (e0 = op(e3, e6))) | ~ sP27), inference(usedef, [], [e33])).
fof(e33, plain, (sP27 <=> (~ (e6 = op(e3, e6)) & ~ (e5 = op(e3, e6)) & ~ (e4 = op(e3, e6)) & ~ (e3 = op(e3, e6)) & ~ (e2 = op(e3, e6)) & ~ (e1 = op(e3, e6)) & ~ (e0 = op(e3, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f2104, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2103, f1873])).
fof(f1873, plain, ~ sP28, inference(subsumption_resolution, [], [f1194, f348])).
fof(f1194, plain, (~ (e6 = op(e4, e0)) | ~ sP28), inference(cnf_transformation, [], [f270])).
fof(f270, plain, ((~ (e6 = op(e4, e0)) & ~ (e5 = op(e4, e0)) & ~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP28), inference(nnf_transformation, [], [f34])).
fof(f34, plain, ((~ (e6 = op(e4, e0)) & ~ (e5 = op(e4, e0)) & ~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP28), inference(usedef, [], [e34])).
fof(e34, plain, (sP28 <=> (~ (e6 = op(e4, e0)) & ~ (e5 = op(e4, e0)) & ~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f2103, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2102, f1872])).
fof(f1872, plain, ~ sP29, inference(subsumption_resolution, [], [f1183, f349])).
fof(f1183, plain, (~ (e2 = op(e4, e1)) | ~ sP29), inference(cnf_transformation, [], [f269])).
fof(f269, plain, ((~ (e6 = op(e4, e1)) & ~ (e5 = op(e4, e1)) & ~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP29), inference(nnf_transformation, [], [f35])).
fof(f35, plain, ((~ (e6 = op(e4, e1)) & ~ (e5 = op(e4, e1)) & ~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP29), inference(usedef, [], [e35])).
fof(e35, plain, (sP29 <=> (~ (e6 = op(e4, e1)) & ~ (e5 = op(e4, e1)) & ~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f2102, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2101, f1867])).
fof(f1867, plain, ~ sP30, inference(subsumption_resolution, [], [f1179, f350])).
fof(f1179, plain, (~ (e5 = op(e4, e2)) | ~ sP30), inference(cnf_transformation, [], [f268])).
fof(f268, plain, ((~ (e6 = op(e4, e2)) & ~ (e5 = op(e4, e2)) & ~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP30), inference(nnf_transformation, [], [f36])).
fof(f36, plain, ((~ (e6 = op(e4, e2)) & ~ (e5 = op(e4, e2)) & ~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP30), inference(usedef, [], [e36])).
fof(e36, plain, (sP30 <=> (~ (e6 = op(e4, e2)) & ~ (e5 = op(e4, e2)) & ~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f2101, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2100, f1865])).
fof(f1865, plain, ~ sP31, inference(subsumption_resolution, [], [f1167, f351])).
fof(f1167, plain, (~ (e0 = op(e4, e3)) | ~ sP31), inference(cnf_transformation, [], [f267])).
fof(f267, plain, ((~ (e6 = op(e4, e3)) & ~ (e5 = op(e4, e3)) & ~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP31), inference(nnf_transformation, [], [f37])).
fof(f37, plain, ((~ (e6 = op(e4, e3)) & ~ (e5 = op(e4, e3)) & ~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP31), inference(usedef, [], [e37])).
fof(e37, plain, (sP31 <=> (~ (e6 = op(e4, e3)) & ~ (e5 = op(e4, e3)) & ~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f2100, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2099, f1858])).
fof(f1858, plain, ~ sP32, inference(subsumption_resolution, [], [f1164, f352])).
fof(f1164, plain, (~ (e4 = op(e4, e4)) | ~ sP32), inference(cnf_transformation, [], [f266])).
fof(f266, plain, ((~ (e6 = op(e4, e4)) & ~ (e5 = op(e4, e4)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP32), inference(nnf_transformation, [], [f38])).
fof(f38, plain, ((~ (e6 = op(e4, e4)) & ~ (e5 = op(e4, e4)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP32), inference(usedef, [], [e38])).
fof(e38, plain, (sP32 <=> (~ (e6 = op(e4, e4)) & ~ (e5 = op(e4, e4)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f2099, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2098, f1855])).
fof(f1855, plain, ~ sP33, inference(subsumption_resolution, [], [f1154, f353])).
fof(f1154, plain, (~ (e1 = op(e4, e5)) | ~ sP33), inference(cnf_transformation, [], [f265])).
fof(f265, plain, ((~ (e6 = op(e4, e5)) & ~ (e5 = op(e4, e5)) & ~ (e4 = op(e4, e5)) & ~ (e3 = op(e4, e5)) & ~ (e2 = op(e4, e5)) & ~ (e1 = op(e4, e5)) & ~ (e0 = op(e4, e5))) | ~ sP33), inference(nnf_transformation, [], [f39])).
fof(f39, plain, ((~ (e6 = op(e4, e5)) & ~ (e5 = op(e4, e5)) & ~ (e4 = op(e4, e5)) & ~ (e3 = op(e4, e5)) & ~ (e2 = op(e4, e5)) & ~ (e1 = op(e4, e5)) & ~ (e0 = op(e4, e5))) | ~ sP33), inference(usedef, [], [e39])).
fof(e39, plain, (sP33 <=> (~ (e6 = op(e4, e5)) & ~ (e5 = op(e4, e5)) & ~ (e4 = op(e4, e5)) & ~ (e3 = op(e4, e5)) & ~ (e2 = op(e4, e5)) & ~ (e1 = op(e4, e5)) & ~ (e0 = op(e4, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f2098, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2097, f1849])).
fof(f1849, plain, ~ sP34, inference(subsumption_resolution, [], [f1149, f354])).
fof(f1149, plain, (~ (e3 = op(e4, e6)) | ~ sP34), inference(cnf_transformation, [], [f264])).
fof(f264, plain, ((~ (e6 = op(e4, e6)) & ~ (e5 = op(e4, e6)) & ~ (e4 = op(e4, e6)) & ~ (e3 = op(e4, e6)) & ~ (e2 = op(e4, e6)) & ~ (e1 = op(e4, e6)) & ~ (e0 = op(e4, e6))) | ~ sP34), inference(nnf_transformation, [], [f40])).
fof(f40, plain, ((~ (e6 = op(e4, e6)) & ~ (e5 = op(e4, e6)) & ~ (e4 = op(e4, e6)) & ~ (e3 = op(e4, e6)) & ~ (e2 = op(e4, e6)) & ~ (e1 = op(e4, e6)) & ~ (e0 = op(e4, e6))) | ~ sP34), inference(usedef, [], [e40])).
fof(e40, plain, (sP34 <=> (~ (e6 = op(e4, e6)) & ~ (e5 = op(e4, e6)) & ~ (e4 = op(e4, e6)) & ~ (e3 = op(e4, e6)) & ~ (e2 = op(e4, e6)) & ~ (e1 = op(e4, e6)) & ~ (e0 = op(e4, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f2097, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2096, f1845])).
fof(f1845, plain, ~ sP35, inference(subsumption_resolution, [], [f1143, f355])).
fof(f1143, plain, (~ (e4 = op(e5, e0)) | ~ sP35), inference(cnf_transformation, [], [f263])).
fof(f263, plain, ((~ (e6 = op(e5, e0)) & ~ (e5 = op(e5, e0)) & ~ (e4 = op(e5, e0)) & ~ (e3 = op(e5, e0)) & ~ (e2 = op(e5, e0)) & ~ (e1 = op(e5, e0)) & ~ (e0 = op(e5, e0))) | ~ sP35), inference(nnf_transformation, [], [f41])).
fof(f41, plain, ((~ (e6 = op(e5, e0)) & ~ (e5 = op(e5, e0)) & ~ (e4 = op(e5, e0)) & ~ (e3 = op(e5, e0)) & ~ (e2 = op(e5, e0)) & ~ (e1 = op(e5, e0)) & ~ (e0 = op(e5, e0))) | ~ sP35), inference(usedef, [], [e41])).
fof(e41, plain, (sP35 <=> (~ (e6 = op(e5, e0)) & ~ (e5 = op(e5, e0)) & ~ (e4 = op(e5, e0)) & ~ (e3 = op(e5, e0)) & ~ (e2 = op(e5, e0)) & ~ (e1 = op(e5, e0)) & ~ (e0 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f2096, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2095, f1842])).
fof(f1842, plain, ~ sP36, inference(subsumption_resolution, [], [f1132, f356])).
fof(f1132, plain, (~ (e0 = op(e5, e1)) | ~ sP36), inference(cnf_transformation, [], [f262])).
fof(f262, plain, ((~ (e6 = op(e5, e1)) & ~ (e5 = op(e5, e1)) & ~ (e4 = op(e5, e1)) & ~ (e3 = op(e5, e1)) & ~ (e2 = op(e5, e1)) & ~ (e1 = op(e5, e1)) & ~ (e0 = op(e5, e1))) | ~ sP36), inference(nnf_transformation, [], [f42])).
fof(f42, plain, ((~ (e6 = op(e5, e1)) & ~ (e5 = op(e5, e1)) & ~ (e4 = op(e5, e1)) & ~ (e3 = op(e5, e1)) & ~ (e2 = op(e5, e1)) & ~ (e1 = op(e5, e1)) & ~ (e0 = op(e5, e1))) | ~ sP36), inference(usedef, [], [e42])).
fof(e42, plain, (sP36 <=> (~ (e6 = op(e5, e1)) & ~ (e5 = op(e5, e1)) & ~ (e4 = op(e5, e1)) & ~ (e3 = op(e5, e1)) & ~ (e2 = op(e5, e1)) & ~ (e1 = op(e5, e1)) & ~ (e0 = op(e5, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f2095, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2094, f1835])).
fof(f1835, plain, ~ sP37, inference(subsumption_resolution, [], [f1128, f357])).
fof(f1128, plain, (~ (e3 = op(e5, e2)) | ~ sP37), inference(cnf_transformation, [], [f261])).
fof(f261, plain, ((~ (e6 = op(e5, e2)) & ~ (e5 = op(e5, e2)) & ~ (e4 = op(e5, e2)) & ~ (e3 = op(e5, e2)) & ~ (e2 = op(e5, e2)) & ~ (e1 = op(e5, e2)) & ~ (e0 = op(e5, e2))) | ~ sP37), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ((~ (e6 = op(e5, e2)) & ~ (e5 = op(e5, e2)) & ~ (e4 = op(e5, e2)) & ~ (e3 = op(e5, e2)) & ~ (e2 = op(e5, e2)) & ~ (e1 = op(e5, e2)) & ~ (e0 = op(e5, e2))) | ~ sP37), inference(usedef, [], [e43])).
fof(e43, plain, (sP37 <=> (~ (e6 = op(e5, e2)) & ~ (e5 = op(e5, e2)) & ~ (e4 = op(e5, e2)) & ~ (e3 = op(e5, e2)) & ~ (e2 = op(e5, e2)) & ~ (e1 = op(e5, e2)) & ~ (e0 = op(e5, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f2094, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2093, f1831])).
fof(f1831, plain, ~ sP38, inference(subsumption_resolution, [], [f1124, f358])).
fof(f1124, plain, (~ (e6 = op(e5, e3)) | ~ sP38), inference(cnf_transformation, [], [f260])).
fof(f260, plain, ((~ (e6 = op(e5, e3)) & ~ (e5 = op(e5, e3)) & ~ (e4 = op(e5, e3)) & ~ (e3 = op(e5, e3)) & ~ (e2 = op(e5, e3)) & ~ (e1 = op(e5, e3)) & ~ (e0 = op(e5, e3))) | ~ sP38), inference(nnf_transformation, [], [f44])).
fof(f44, plain, ((~ (e6 = op(e5, e3)) & ~ (e5 = op(e5, e3)) & ~ (e4 = op(e5, e3)) & ~ (e3 = op(e5, e3)) & ~ (e2 = op(e5, e3)) & ~ (e1 = op(e5, e3)) & ~ (e0 = op(e5, e3))) | ~ sP38), inference(usedef, [], [e44])).
fof(e44, plain, (sP38 <=> (~ (e6 = op(e5, e3)) & ~ (e5 = op(e5, e3)) & ~ (e4 = op(e5, e3)) & ~ (e3 = op(e5, e3)) & ~ (e2 = op(e5, e3)) & ~ (e1 = op(e5, e3)) & ~ (e0 = op(e5, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f2093, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2092, f1830])).
fof(f1830, plain, ~ sP39, inference(subsumption_resolution, [], [f1112, f359])).
fof(f1112, plain, (~ (e1 = op(e5, e4)) | ~ sP39), inference(cnf_transformation, [], [f259])).
fof(f259, plain, ((~ (e6 = op(e5, e4)) & ~ (e5 = op(e5, e4)) & ~ (e4 = op(e5, e4)) & ~ (e3 = op(e5, e4)) & ~ (e2 = op(e5, e4)) & ~ (e1 = op(e5, e4)) & ~ (e0 = op(e5, e4))) | ~ sP39), inference(nnf_transformation, [], [f45])).
fof(f45, plain, ((~ (e6 = op(e5, e4)) & ~ (e5 = op(e5, e4)) & ~ (e4 = op(e5, e4)) & ~ (e3 = op(e5, e4)) & ~ (e2 = op(e5, e4)) & ~ (e1 = op(e5, e4)) & ~ (e0 = op(e5, e4))) | ~ sP39), inference(usedef, [], [e45])).
fof(e45, plain, (sP39 <=> (~ (e6 = op(e5, e4)) & ~ (e5 = op(e5, e4)) & ~ (e4 = op(e5, e4)) & ~ (e3 = op(e5, e4)) & ~ (e2 = op(e5, e4)) & ~ (e1 = op(e5, e4)) & ~ (e0 = op(e5, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f2092, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2091, f1824])).
fof(f1824, plain, ~ sP40, inference(subsumption_resolution, [], [f1109, f360])).
fof(f1109, plain, (~ (e5 = op(e5, e5)) | ~ sP40), inference(cnf_transformation, [], [f258])).
fof(f258, plain, ((~ (e6 = op(e5, e5)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e5, e5)) & ~ (e3 = op(e5, e5)) & ~ (e2 = op(e5, e5)) & ~ (e1 = op(e5, e5)) & ~ (e0 = op(e5, e5))) | ~ sP40), inference(nnf_transformation, [], [f46])).
fof(f46, plain, ((~ (e6 = op(e5, e5)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e5, e5)) & ~ (e3 = op(e5, e5)) & ~ (e2 = op(e5, e5)) & ~ (e1 = op(e5, e5)) & ~ (e0 = op(e5, e5))) | ~ sP40), inference(usedef, [], [e46])).
fof(e46, plain, (sP40 <=> (~ (e6 = op(e5, e5)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e5, e5)) & ~ (e3 = op(e5, e5)) & ~ (e2 = op(e5, e5)) & ~ (e1 = op(e5, e5)) & ~ (e0 = op(e5, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f2091, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2090, f1822])).
fof(f1822, plain, ~ sP41, inference(subsumption_resolution, [], [f1099, f361])).
fof(f1099, plain, (~ (e2 = op(e5, e6)) | ~ sP41), inference(cnf_transformation, [], [f257])).
fof(f257, plain, ((~ (e6 = op(e5, e6)) & ~ (e5 = op(e5, e6)) & ~ (e4 = op(e5, e6)) & ~ (e3 = op(e5, e6)) & ~ (e2 = op(e5, e6)) & ~ (e1 = op(e5, e6)) & ~ (e0 = op(e5, e6))) | ~ sP41), inference(nnf_transformation, [], [f47])).
fof(f47, plain, ((~ (e6 = op(e5, e6)) & ~ (e5 = op(e5, e6)) & ~ (e4 = op(e5, e6)) & ~ (e3 = op(e5, e6)) & ~ (e2 = op(e5, e6)) & ~ (e1 = op(e5, e6)) & ~ (e0 = op(e5, e6))) | ~ sP41), inference(usedef, [], [e47])).
fof(e47, plain, (sP41 <=> (~ (e6 = op(e5, e6)) & ~ (e5 = op(e5, e6)) & ~ (e4 = op(e5, e6)) & ~ (e3 = op(e5, e6)) & ~ (e2 = op(e5, e6)) & ~ (e1 = op(e5, e6)) & ~ (e0 = op(e5, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f2090, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2089, f1817])).
fof(f1817, plain, ~ sP42, inference(subsumption_resolution, [], [f1095, f362])).
fof(f1095, plain, (~ (e5 = op(e6, e0)) | ~ sP42), inference(cnf_transformation, [], [f256])).
fof(f256, plain, ((~ (e6 = op(e6, e0)) & ~ (e5 = op(e6, e0)) & ~ (e4 = op(e6, e0)) & ~ (e3 = op(e6, e0)) & ~ (e2 = op(e6, e0)) & ~ (e1 = op(e6, e0)) & ~ (e0 = op(e6, e0))) | ~ sP42), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ((~ (e6 = op(e6, e0)) & ~ (e5 = op(e6, e0)) & ~ (e4 = op(e6, e0)) & ~ (e3 = op(e6, e0)) & ~ (e2 = op(e6, e0)) & ~ (e1 = op(e6, e0)) & ~ (e0 = op(e6, e0))) | ~ sP42), inference(usedef, [], [e48])).
fof(e48, plain, (sP42 <=> (~ (e6 = op(e6, e0)) & ~ (e5 = op(e6, e0)) & ~ (e4 = op(e6, e0)) & ~ (e3 = op(e6, e0)) & ~ (e2 = op(e6, e0)) & ~ (e1 = op(e6, e0)) & ~ (e0 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f2089, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2088, f1815])).
fof(f1815, plain, ~ sP43, inference(subsumption_resolution, [], [f1087, f363])).
fof(f1087, plain, (~ (e4 = op(e6, e1)) | ~ sP43), inference(cnf_transformation, [], [f255])).
fof(f255, plain, ((~ (e6 = op(e6, e1)) & ~ (e5 = op(e6, e1)) & ~ (e4 = op(e6, e1)) & ~ (e3 = op(e6, e1)) & ~ (e2 = op(e6, e1)) & ~ (e1 = op(e6, e1)) & ~ (e0 = op(e6, e1))) | ~ sP43), inference(nnf_transformation, [], [f49])).
fof(f49, plain, ((~ (e6 = op(e6, e1)) & ~ (e5 = op(e6, e1)) & ~ (e4 = op(e6, e1)) & ~ (e3 = op(e6, e1)) & ~ (e2 = op(e6, e1)) & ~ (e1 = op(e6, e1)) & ~ (e0 = op(e6, e1))) | ~ sP43), inference(usedef, [], [e49])).
fof(e49, plain, (sP43 <=> (~ (e6 = op(e6, e1)) & ~ (e5 = op(e6, e1)) & ~ (e4 = op(e6, e1)) & ~ (e3 = op(e6, e1)) & ~ (e2 = op(e6, e1)) & ~ (e1 = op(e6, e1)) & ~ (e0 = op(e6, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f2088, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2087, f1812])).
fof(f1812, plain, ~ sP44, inference(subsumption_resolution, [], [f1076, f364])).
fof(f1076, plain, (~ (e0 = op(e6, e2)) | ~ sP44), inference(cnf_transformation, [], [f254])).
fof(f254, plain, ((~ (e6 = op(e6, e2)) & ~ (e5 = op(e6, e2)) & ~ (e4 = op(e6, e2)) & ~ (e3 = op(e6, e2)) & ~ (e2 = op(e6, e2)) & ~ (e1 = op(e6, e2)) & ~ (e0 = op(e6, e2))) | ~ sP44), inference(nnf_transformation, [], [f50])).
fof(f50, plain, ((~ (e6 = op(e6, e2)) & ~ (e5 = op(e6, e2)) & ~ (e4 = op(e6, e2)) & ~ (e3 = op(e6, e2)) & ~ (e2 = op(e6, e2)) & ~ (e1 = op(e6, e2)) & ~ (e0 = op(e6, e2))) | ~ sP44), inference(usedef, [], [e50])).
fof(e50, plain, (sP44 <=> (~ (e6 = op(e6, e2)) & ~ (e5 = op(e6, e2)) & ~ (e4 = op(e6, e2)) & ~ (e3 = op(e6, e2)) & ~ (e2 = op(e6, e2)) & ~ (e1 = op(e6, e2)) & ~ (e0 = op(e6, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f2087, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2086, f1805])).
fof(f1805, plain, ~ sP45, inference(subsumption_resolution, [], [f1070, f365])).
fof(f1070, plain, (~ (e1 = op(e6, e3)) | ~ sP45), inference(cnf_transformation, [], [f253])).
fof(f253, plain, ((~ (e6 = op(e6, e3)) & ~ (e5 = op(e6, e3)) & ~ (e4 = op(e6, e3)) & ~ (e3 = op(e6, e3)) & ~ (e2 = op(e6, e3)) & ~ (e1 = op(e6, e3)) & ~ (e0 = op(e6, e3))) | ~ sP45), inference(nnf_transformation, [], [f51])).
fof(f51, plain, ((~ (e6 = op(e6, e3)) & ~ (e5 = op(e6, e3)) & ~ (e4 = op(e6, e3)) & ~ (e3 = op(e6, e3)) & ~ (e2 = op(e6, e3)) & ~ (e1 = op(e6, e3)) & ~ (e0 = op(e6, e3))) | ~ sP45), inference(usedef, [], [e51])).
fof(e51, plain, (sP45 <=> (~ (e6 = op(e6, e3)) & ~ (e5 = op(e6, e3)) & ~ (e4 = op(e6, e3)) & ~ (e3 = op(e6, e3)) & ~ (e2 = op(e6, e3)) & ~ (e1 = op(e6, e3)) & ~ (e0 = op(e6, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f2086, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2085, f1799])).
fof(f1799, plain, ~ sP46, inference(subsumption_resolution, [], [f1065, f366])).
fof(f1065, plain, (~ (e3 = op(e6, e4)) | ~ sP46), inference(cnf_transformation, [], [f252])).
fof(f252, plain, ((~ (e6 = op(e6, e4)) & ~ (e5 = op(e6, e4)) & ~ (e4 = op(e6, e4)) & ~ (e3 = op(e6, e4)) & ~ (e2 = op(e6, e4)) & ~ (e1 = op(e6, e4)) & ~ (e0 = op(e6, e4))) | ~ sP46), inference(nnf_transformation, [], [f52])).
fof(f52, plain, ((~ (e6 = op(e6, e4)) & ~ (e5 = op(e6, e4)) & ~ (e4 = op(e6, e4)) & ~ (e3 = op(e6, e4)) & ~ (e2 = op(e6, e4)) & ~ (e1 = op(e6, e4)) & ~ (e0 = op(e6, e4))) | ~ sP46), inference(usedef, [], [e52])).
fof(e52, plain, (sP46 <=> (~ (e6 = op(e6, e4)) & ~ (e5 = op(e6, e4)) & ~ (e4 = op(e6, e4)) & ~ (e3 = op(e6, e4)) & ~ (e2 = op(e6, e4)) & ~ (e1 = op(e6, e4)) & ~ (e0 = op(e6, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f2085, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2084, f1795])).
fof(f1795, plain, ~ sP47, inference(subsumption_resolution, [], [f1057, f367])).
fof(f1057, plain, (~ (e2 = op(e6, e5)) | ~ sP47), inference(cnf_transformation, [], [f251])).
fof(f251, plain, ((~ (e6 = op(e6, e5)) & ~ (e5 = op(e6, e5)) & ~ (e4 = op(e6, e5)) & ~ (e3 = op(e6, e5)) & ~ (e2 = op(e6, e5)) & ~ (e1 = op(e6, e5)) & ~ (e0 = op(e6, e5))) | ~ sP47), inference(nnf_transformation, [], [f53])).
fof(f53, plain, ((~ (e6 = op(e6, e5)) & ~ (e5 = op(e6, e5)) & ~ (e4 = op(e6, e5)) & ~ (e3 = op(e6, e5)) & ~ (e2 = op(e6, e5)) & ~ (e1 = op(e6, e5)) & ~ (e0 = op(e6, e5))) | ~ sP47), inference(usedef, [], [e53])).
fof(e53, plain, (sP47 <=> (~ (e6 = op(e6, e5)) & ~ (e5 = op(e6, e5)) & ~ (e4 = op(e6, e5)) & ~ (e3 = op(e6, e5)) & ~ (e2 = op(e6, e5)) & ~ (e1 = op(e6, e5)) & ~ (e0 = op(e6, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f2084, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2083, f1790])).
fof(f1790, plain, ~ sP48, inference(subsumption_resolution, [], [f1054, f368])).
fof(f1054, plain, (~ (e6 = op(e6, e6)) | ~ sP48), inference(cnf_transformation, [], [f250])).
fof(f250, plain, ((~ (e6 = op(e6, e6)) & ~ (e5 = op(e6, e6)) & ~ (e4 = op(e6, e6)) & ~ (e3 = op(e6, e6)) & ~ (e2 = op(e6, e6)) & ~ (e1 = op(e6, e6)) & ~ (e0 = op(e6, e6))) | ~ sP48), inference(nnf_transformation, [], [f54])).
fof(f54, plain, ((~ (e6 = op(e6, e6)) & ~ (e5 = op(e6, e6)) & ~ (e4 = op(e6, e6)) & ~ (e3 = op(e6, e6)) & ~ (e2 = op(e6, e6)) & ~ (e1 = op(e6, e6)) & ~ (e0 = op(e6, e6))) | ~ sP48), inference(usedef, [], [e54])).
fof(e54, plain, (sP48 <=> (~ (e6 = op(e6, e6)) & ~ (e5 = op(e6, e6)) & ~ (e4 = op(e6, e6)) & ~ (e3 = op(e6, e6)) & ~ (e2 = op(e6, e6)) & ~ (e1 = op(e6, e6)) & ~ (e0 = op(e6, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f2083, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2082, f1789])).
fof(f1789, plain, ~ sP49, inference(subsumption_resolution, [], [f1041, f320])).
fof(f1041, plain, (~ (e0 = op(e0, e0)) | ~ sP49), inference(cnf_transformation, [], [f249])).
fof(f249, plain, ((~ (e0 = op(e0, e6)) & ~ (e0 = op(e0, e5)) & ~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP49), inference(nnf_transformation, [], [f55])).
fof(f55, plain, ((~ (e0 = op(e0, e6)) & ~ (e0 = op(e0, e5)) & ~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP49), inference(usedef, [], [e55])).
fof(e55, plain, (sP49 <=> (~ (e0 = op(e0, e6)) & ~ (e0 = op(e0, e5)) & ~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f2082, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2081, f1782])).
fof(f1782, plain, ~ sP50, inference(subsumption_resolution, [], [f1034, f320])).
fof(f1034, plain, (~ (e0 = op(e0, e0)) | ~ sP50), inference(cnf_transformation, [], [f248])).
fof(f248, plain, ((~ (e0 = op(e6, e0)) & ~ (e0 = op(e5, e0)) & ~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP50), inference(nnf_transformation, [], [f56])).
fof(f56, plain, ((~ (e0 = op(e6, e0)) & ~ (e0 = op(e5, e0)) & ~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP50), inference(usedef, [], [e56])).
fof(e56, plain, (sP50 <=> (~ (e0 = op(e6, e0)) & ~ (e0 = op(e5, e0)) & ~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f2081, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2080, f1775])).
fof(f1775, plain, ~ sP51, inference(subsumption_resolution, [], [f1029, f322])).
fof(f1029, plain, (~ (e1 = op(e0, e2)) | ~ sP51), inference(cnf_transformation, [], [f247])).
fof(f247, plain, ((~ (e1 = op(e0, e6)) & ~ (e1 = op(e0, e5)) & ~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP51), inference(nnf_transformation, [], [f57])).
fof(f57, plain, ((~ (e1 = op(e0, e6)) & ~ (e1 = op(e0, e5)) & ~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP51), inference(usedef, [], [e57])).
fof(e57, plain, (sP51 <=> (~ (e1 = op(e0, e6)) & ~ (e1 = op(e0, e5)) & ~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f2080, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2079, f1770])).
fof(f1770, plain, ~ sP52, inference(subsumption_resolution, [], [f1022, f334])).
fof(f1022, plain, (~ (e1 = op(e2, e0)) | ~ sP52), inference(cnf_transformation, [], [f246])).
fof(f246, plain, ((~ (e1 = op(e6, e0)) & ~ (e1 = op(e5, e0)) & ~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP52), inference(nnf_transformation, [], [f58])).
fof(f58, plain, ((~ (e1 = op(e6, e0)) & ~ (e1 = op(e5, e0)) & ~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP52), inference(usedef, [], [e58])).
fof(e58, plain, (sP52 <=> (~ (e1 = op(e6, e0)) & ~ (e1 = op(e5, e0)) & ~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f2079, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2078, f1765])).
fof(f1765, plain, ~ sP53, inference(subsumption_resolution, [], [f1016, f323])).
fof(f1016, plain, (~ (e2 = op(e0, e3)) | ~ sP53), inference(cnf_transformation, [], [f245])).
fof(f245, plain, ((~ (e2 = op(e0, e6)) & ~ (e2 = op(e0, e5)) & ~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP53), inference(nnf_transformation, [], [f59])).
fof(f59, plain, ((~ (e2 = op(e0, e6)) & ~ (e2 = op(e0, e5)) & ~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP53), inference(usedef, [], [e59])).
fof(e59, plain, (sP53 <=> (~ (e2 = op(e0, e6)) & ~ (e2 = op(e0, e5)) & ~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f2078, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2077, f1761])).
fof(f1761, plain, ~ sP54, inference(subsumption_resolution, [], [f1009, f341])).
fof(f1009, plain, (~ (e2 = op(e3, e0)) | ~ sP54), inference(cnf_transformation, [], [f244])).
fof(f244, plain, ((~ (e2 = op(e6, e0)) & ~ (e2 = op(e5, e0)) & ~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP54), inference(nnf_transformation, [], [f60])).
fof(f60, plain, ((~ (e2 = op(e6, e0)) & ~ (e2 = op(e5, e0)) & ~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP54), inference(usedef, [], [e60])).
fof(e60, plain, (sP54 <=> (~ (e2 = op(e6, e0)) & ~ (e2 = op(e5, e0)) & ~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f2077, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2076, f1757])).
fof(f1757, plain, ~ sP55, inference(subsumption_resolution, [], [f1000, f321])).
fof(f1000, plain, (~ (e3 = op(e0, e1)) | ~ sP55), inference(cnf_transformation, [], [f243])).
fof(f243, plain, ((~ (e3 = op(e0, e6)) & ~ (e3 = op(e0, e5)) & ~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP55), inference(nnf_transformation, [], [f61])).
fof(f61, plain, ((~ (e3 = op(e0, e6)) & ~ (e3 = op(e0, e5)) & ~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP55), inference(usedef, [], [e61])).
fof(e61, plain, (sP55 <=> (~ (e3 = op(e0, e6)) & ~ (e3 = op(e0, e5)) & ~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f2076, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2075, f1751])).
fof(f1751, plain, ~ sP56, inference(subsumption_resolution, [], [f993, f327])).
fof(f993, plain, (~ (e3 = op(e1, e0)) | ~ sP56), inference(cnf_transformation, [], [f242])).
fof(f242, plain, ((~ (e3 = op(e6, e0)) & ~ (e3 = op(e5, e0)) & ~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP56), inference(nnf_transformation, [], [f62])).
fof(f62, plain, ((~ (e3 = op(e6, e0)) & ~ (e3 = op(e5, e0)) & ~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP56), inference(usedef, [], [e62])).
fof(e62, plain, (sP56 <=> (~ (e3 = op(e6, e0)) & ~ (e3 = op(e5, e0)) & ~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f2075, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2074, f1745])).
fof(f1745, plain, ~ sP57, inference(subsumption_resolution, [], [f990, f325])).
fof(f990, plain, (~ (e4 = op(e0, e5)) | ~ sP57), inference(cnf_transformation, [], [f241])).
fof(f241, plain, ((~ (e4 = op(e0, e6)) & ~ (e4 = op(e0, e5)) & ~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP57), inference(nnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e4 = op(e0, e6)) & ~ (e4 = op(e0, e5)) & ~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP57), inference(usedef, [], [e63])).
fof(e63, plain, (sP57 <=> (~ (e4 = op(e0, e6)) & ~ (e4 = op(e0, e5)) & ~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f2074, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2073, f1743])).
fof(f1743, plain, ~ sP58, inference(subsumption_resolution, [], [f983, f355])).
fof(f983, plain, (~ (e4 = op(e5, e0)) | ~ sP58), inference(cnf_transformation, [], [f240])).
fof(f240, plain, ((~ (e4 = op(e6, e0)) & ~ (e4 = op(e5, e0)) & ~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP58), inference(nnf_transformation, [], [f64])).
fof(f64, plain, ((~ (e4 = op(e6, e0)) & ~ (e4 = op(e5, e0)) & ~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP58), inference(usedef, [], [e64])).
fof(e64, plain, (sP58 <=> (~ (e4 = op(e6, e0)) & ~ (e4 = op(e5, e0)) & ~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f2073, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2072, f1741])).
fof(f1741, plain, ~ sP59, inference(subsumption_resolution, [], [f977, f326])).
fof(f977, plain, (~ (e5 = op(e0, e6)) | ~ sP59), inference(cnf_transformation, [], [f239])).
fof(f239, plain, ((~ (e5 = op(e0, e6)) & ~ (e5 = op(e0, e5)) & ~ (e5 = op(e0, e4)) & ~ (e5 = op(e0, e3)) & ~ (e5 = op(e0, e2)) & ~ (e5 = op(e0, e1)) & ~ (e5 = op(e0, e0))) | ~ sP59), inference(nnf_transformation, [], [f65])).
fof(f65, plain, ((~ (e5 = op(e0, e6)) & ~ (e5 = op(e0, e5)) & ~ (e5 = op(e0, e4)) & ~ (e5 = op(e0, e3)) & ~ (e5 = op(e0, e2)) & ~ (e5 = op(e0, e1)) & ~ (e5 = op(e0, e0))) | ~ sP59), inference(usedef, [], [e65])).
fof(e65, plain, (sP59 <=> (~ (e5 = op(e0, e6)) & ~ (e5 = op(e0, e5)) & ~ (e5 = op(e0, e4)) & ~ (e5 = op(e0, e3)) & ~ (e5 = op(e0, e2)) & ~ (e5 = op(e0, e1)) & ~ (e5 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f2072, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2071, f1740])).
fof(f1740, plain, ~ sP60, inference(subsumption_resolution, [], [f970, f362])).
fof(f970, plain, (~ (e5 = op(e6, e0)) | ~ sP60), inference(cnf_transformation, [], [f238])).
fof(f238, plain, ((~ (e5 = op(e6, e0)) & ~ (e5 = op(e5, e0)) & ~ (e5 = op(e4, e0)) & ~ (e5 = op(e3, e0)) & ~ (e5 = op(e2, e0)) & ~ (e5 = op(e1, e0)) & ~ (e5 = op(e0, e0))) | ~ sP60), inference(nnf_transformation, [], [f66])).
fof(f66, plain, ((~ (e5 = op(e6, e0)) & ~ (e5 = op(e5, e0)) & ~ (e5 = op(e4, e0)) & ~ (e5 = op(e3, e0)) & ~ (e5 = op(e2, e0)) & ~ (e5 = op(e1, e0)) & ~ (e5 = op(e0, e0))) | ~ sP60), inference(usedef, [], [e66])).
fof(e66, plain, (sP60 <=> (~ (e5 = op(e6, e0)) & ~ (e5 = op(e5, e0)) & ~ (e5 = op(e4, e0)) & ~ (e5 = op(e3, e0)) & ~ (e5 = op(e2, e0)) & ~ (e5 = op(e1, e0)) & ~ (e5 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f2071, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2070, f1739])).
fof(f1739, plain, ~ sP61, inference(subsumption_resolution, [], [f961, f324])).
fof(f961, plain, (~ (e6 = op(e0, e4)) | ~ sP61), inference(cnf_transformation, [], [f237])).
fof(f237, plain, ((~ (e6 = op(e0, e6)) & ~ (e6 = op(e0, e5)) & ~ (e6 = op(e0, e4)) & ~ (e6 = op(e0, e3)) & ~ (e6 = op(e0, e2)) & ~ (e6 = op(e0, e1)) & ~ (e6 = op(e0, e0))) | ~ sP61), inference(nnf_transformation, [], [f67])).
fof(f67, plain, ((~ (e6 = op(e0, e6)) & ~ (e6 = op(e0, e5)) & ~ (e6 = op(e0, e4)) & ~ (e6 = op(e0, e3)) & ~ (e6 = op(e0, e2)) & ~ (e6 = op(e0, e1)) & ~ (e6 = op(e0, e0))) | ~ sP61), inference(usedef, [], [e67])).
fof(e67, plain, (sP61 <=> (~ (e6 = op(e0, e6)) & ~ (e6 = op(e0, e5)) & ~ (e6 = op(e0, e4)) & ~ (e6 = op(e0, e3)) & ~ (e6 = op(e0, e2)) & ~ (e6 = op(e0, e1)) & ~ (e6 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f2070, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2069, f1736])).
fof(f1736, plain, ~ sP62, inference(subsumption_resolution, [], [f954, f348])).
fof(f954, plain, (~ (e6 = op(e4, e0)) | ~ sP62), inference(cnf_transformation, [], [f236])).
fof(f236, plain, ((~ (e6 = op(e6, e0)) & ~ (e6 = op(e5, e0)) & ~ (e6 = op(e4, e0)) & ~ (e6 = op(e3, e0)) & ~ (e6 = op(e2, e0)) & ~ (e6 = op(e1, e0)) & ~ (e6 = op(e0, e0))) | ~ sP62), inference(nnf_transformation, [], [f68])).
fof(f68, plain, ((~ (e6 = op(e6, e0)) & ~ (e6 = op(e5, e0)) & ~ (e6 = op(e4, e0)) & ~ (e6 = op(e3, e0)) & ~ (e6 = op(e2, e0)) & ~ (e6 = op(e1, e0)) & ~ (e6 = op(e0, e0))) | ~ sP62), inference(usedef, [], [e68])).
fof(e68, plain, (sP62 <=> (~ (e6 = op(e6, e0)) & ~ (e6 = op(e5, e0)) & ~ (e6 = op(e4, e0)) & ~ (e6 = op(e3, e0)) & ~ (e6 = op(e2, e0)) & ~ (e6 = op(e1, e0)) & ~ (e6 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f2069, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2068, f1733])).
fof(f1733, plain, ~ sP63, inference(subsumption_resolution, [], [f948, f332])).
fof(f948, plain, (~ (e0 = op(e1, e5)) | ~ sP63), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ((~ (e0 = op(e1, e6)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP63), inference(nnf_transformation, [], [f69])).
fof(f69, plain, ((~ (e0 = op(e1, e6)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP63), inference(usedef, [], [e69])).
fof(e69, plain, (sP63 <=> (~ (e0 = op(e1, e6)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f2068, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2067, f1731])).
fof(f1731, plain, ~ sP64, inference(subsumption_resolution, [], [f941, f356])).
fof(f941, plain, (~ (e0 = op(e5, e1)) | ~ sP64), inference(cnf_transformation, [], [f234])).
fof(f234, plain, ((~ (e0 = op(e6, e1)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP64), inference(nnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e0 = op(e6, e1)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP64), inference(usedef, [], [e70])).
fof(e70, plain, (sP64 <=> (~ (e0 = op(e6, e1)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f2067, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2066, f1729])).
fof(f1729, plain, ~ sP65, inference(subsumption_resolution, [], [f930, f328])).
fof(f930, plain, (~ (e1 = op(e1, e1)) | ~ sP65), inference(cnf_transformation, [], [f233])).
fof(f233, plain, ((~ (e1 = op(e1, e6)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP65), inference(nnf_transformation, [], [f71])).
fof(f71, plain, ((~ (e1 = op(e1, e6)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP65), inference(usedef, [], [e71])).
fof(e71, plain, (sP65 <=> (~ (e1 = op(e1, e6)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f2066, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2065, f1723])).
fof(f1723, plain, ~ sP66, inference(subsumption_resolution, [], [f923, f328])).
fof(f923, plain, (~ (e1 = op(e1, e1)) | ~ sP66), inference(cnf_transformation, [], [f232])).
fof(f232, plain, ((~ (e1 = op(e6, e1)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP66), inference(nnf_transformation, [], [f72])).
fof(f72, plain, ((~ (e1 = op(e6, e1)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP66), inference(usedef, [], [e72])).
fof(e72, plain, (sP66 <=> (~ (e1 = op(e6, e1)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f2065, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2064, f1717])).
fof(f1717, plain, ~ sP67, inference(subsumption_resolution, [], [f919, f331])).
fof(f919, plain, (~ (e2 = op(e1, e4)) | ~ sP67), inference(cnf_transformation, [], [f231])).
fof(f231, plain, ((~ (e2 = op(e1, e6)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP67), inference(nnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e2 = op(e1, e6)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP67), inference(usedef, [], [e73])).
fof(e73, plain, (sP67 <=> (~ (e2 = op(e1, e6)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f2064, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2063, f1714])).
fof(f1714, plain, ~ sP68, inference(subsumption_resolution, [], [f912, f349])).
fof(f912, plain, (~ (e2 = op(e4, e1)) | ~ sP68), inference(cnf_transformation, [], [f230])).
fof(f230, plain, ((~ (e2 = op(e6, e1)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP68), inference(nnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e2 = op(e6, e1)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP68), inference(usedef, [], [e74])).
fof(e74, plain, (sP68 <=> (~ (e2 = op(e6, e1)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f2063, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2062, f1711])).
fof(f1711, plain, ~ sP69, inference(subsumption_resolution, [], [f901, f327])).
fof(f901, plain, (~ (e3 = op(e1, e0)) | ~ sP69), inference(cnf_transformation, [], [f229])).
fof(f229, plain, ((~ (e3 = op(e1, e6)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP69), inference(nnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e3 = op(e1, e6)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP69), inference(usedef, [], [e75])).
fof(e75, plain, (sP69 <=> (~ (e3 = op(e1, e6)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f2062, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2061, f1704])).
fof(f1704, plain, ~ sP70, inference(subsumption_resolution, [], [f894, f321])).
fof(f894, plain, (~ (e3 = op(e0, e1)) | ~ sP70), inference(cnf_transformation, [], [f228])).
fof(f228, plain, ((~ (e3 = op(e6, e1)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP70), inference(nnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e3 = op(e6, e1)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP70), inference(usedef, [], [e76])).
fof(e76, plain, (sP70 <=> (~ (e3 = op(e6, e1)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f2061, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2060, f1697])).
fof(f1697, plain, ~ sP71, inference(subsumption_resolution, [], [f893, f333])).
fof(f893, plain, (~ (e4 = op(e1, e6)) | ~ sP71), inference(cnf_transformation, [], [f227])).
fof(f227, plain, ((~ (e4 = op(e1, e6)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP71), inference(nnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e4 = op(e1, e6)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP71), inference(usedef, [], [e77])).
fof(e77, plain, (sP71 <=> (~ (e4 = op(e1, e6)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f2060, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2059, f1696])).
fof(f1696, plain, ~ sP72, inference(subsumption_resolution, [], [f886, f363])).
fof(f886, plain, (~ (e4 = op(e6, e1)) | ~ sP72), inference(cnf_transformation, [], [f226])).
fof(f226, plain, ((~ (e4 = op(e6, e1)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP72), inference(nnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e4 = op(e6, e1)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP72), inference(usedef, [], [e78])).
fof(e78, plain, (sP72 <=> (~ (e4 = op(e6, e1)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f2059, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2058, f1695])).
fof(f1695, plain, ~ sP73, inference(subsumption_resolution, [], [f876, f330])).
fof(f876, plain, (~ (e5 = op(e1, e3)) | ~ sP73), inference(cnf_transformation, [], [f225])).
fof(f225, plain, ((~ (e5 = op(e1, e6)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e1, e0))) | ~ sP73), inference(nnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e5 = op(e1, e6)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e1, e0))) | ~ sP73), inference(usedef, [], [e79])).
fof(e79, plain, (sP73 <=> (~ (e5 = op(e1, e6)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f2058, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2057, f1691])).
fof(f1691, plain, ~ sP74, inference(subsumption_resolution, [], [f869, f342])).
fof(f869, plain, (~ (e5 = op(e3, e1)) | ~ sP74), inference(cnf_transformation, [], [f224])).
fof(f224, plain, ((~ (e5 = op(e6, e1)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e0, e1))) | ~ sP74), inference(nnf_transformation, [], [f80])).
fof(f80, plain, ((~ (e5 = op(e6, e1)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e0, e1))) | ~ sP74), inference(usedef, [], [e80])).
fof(e80, plain, (sP74 <=> (~ (e5 = op(e6, e1)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f2057, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2056, f1687])).
fof(f1687, plain, ~ sP75, inference(subsumption_resolution, [], [f861, f329])).
fof(f861, plain, (~ (e6 = op(e1, e2)) | ~ sP75), inference(cnf_transformation, [], [f223])).
fof(f223, plain, ((~ (e6 = op(e1, e6)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e1, e0))) | ~ sP75), inference(nnf_transformation, [], [f81])).
fof(f81, plain, ((~ (e6 = op(e1, e6)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e1, e0))) | ~ sP75), inference(usedef, [], [e81])).
fof(e81, plain, (sP75 <=> (~ (e6 = op(e1, e6)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f2056, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2055, f1682])).
fof(f1682, plain, ~ sP76, inference(subsumption_resolution, [], [f854, f335])).
fof(f854, plain, (~ (e6 = op(e2, e1)) | ~ sP76), inference(cnf_transformation, [], [f222])).
fof(f222, plain, ((~ (e6 = op(e6, e1)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e0, e1))) | ~ sP76), inference(nnf_transformation, [], [f82])).
fof(f82, plain, ((~ (e6 = op(e6, e1)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e0, e1))) | ~ sP76), inference(usedef, [], [e82])).
fof(e82, plain, (sP76 <=> (~ (e6 = op(e6, e1)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f2055, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2054, f1677])).
fof(f1677, plain, ~ sP77, inference(subsumption_resolution, [], [f851, f340])).
fof(f851, plain, (~ (e0 = op(e2, e6)) | ~ sP77), inference(cnf_transformation, [], [f221])).
fof(f221, plain, ((~ (e0 = op(e2, e6)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP77), inference(nnf_transformation, [], [f83])).
fof(f83, plain, ((~ (e0 = op(e2, e6)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP77), inference(usedef, [], [e83])).
fof(e83, plain, (sP77 <=> (~ (e0 = op(e2, e6)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f2054, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2053, f1676])).
fof(f1676, plain, ~ sP78, inference(subsumption_resolution, [], [f844, f364])).
fof(f844, plain, (~ (e0 = op(e6, e2)) | ~ sP78), inference(cnf_transformation, [], [f220])).
fof(f220, plain, ((~ (e0 = op(e6, e2)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP78), inference(nnf_transformation, [], [f84])).
fof(f84, plain, ((~ (e0 = op(e6, e2)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP78), inference(usedef, [], [e84])).
fof(e84, plain, (sP78 <=> (~ (e0 = op(e6, e2)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f2053, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2052, f1675])).
fof(f1675, plain, ~ sP79, inference(subsumption_resolution, [], [f831, f334])).
fof(f831, plain, (~ (e1 = op(e2, e0)) | ~ sP79), inference(cnf_transformation, [], [f219])).
fof(f219, plain, ((~ (e1 = op(e2, e6)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP79), inference(nnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e1 = op(e2, e6)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP79), inference(usedef, [], [e85])).
fof(e85, plain, (sP79 <=> (~ (e1 = op(e2, e6)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f2052, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2051, f1668])).
fof(f1668, plain, ~ sP80, inference(subsumption_resolution, [], [f824, f322])).
fof(f824, plain, (~ (e1 = op(e0, e2)) | ~ sP80), inference(cnf_transformation, [], [f218])).
fof(f218, plain, ((~ (e1 = op(e6, e2)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP80), inference(nnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e1 = op(e6, e2)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP80), inference(usedef, [], [e86])).
fof(e86, plain, (sP80 <=> (~ (e1 = op(e6, e2)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f2051, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2050, f1661])).
fof(f1661, plain, ~ sP81, inference(subsumption_resolution, [], [f819, f336])).
fof(f819, plain, (~ (e2 = op(e2, e2)) | ~ sP81), inference(cnf_transformation, [], [f217])).
fof(f217, plain, ((~ (e2 = op(e2, e6)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP81), inference(nnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e2 = op(e2, e6)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP81), inference(usedef, [], [e87])).
fof(e87, plain, (sP81 <=> (~ (e2 = op(e2, e6)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f2050, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2049, f1656])).
fof(f1656, plain, ~ sP82, inference(subsumption_resolution, [], [f812, f336])).
fof(f812, plain, (~ (e2 = op(e2, e2)) | ~ sP82), inference(cnf_transformation, [], [f216])).
fof(f216, plain, ((~ (e2 = op(e6, e2)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP82), inference(nnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e2 = op(e6, e2)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP82), inference(usedef, [], [e88])).
fof(e88, plain, (sP82 <=> (~ (e2 = op(e6, e2)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f2049, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2048, f1651])).
fof(f1651, plain, ~ sP83, inference(subsumption_resolution, [], [f808, f339])).
fof(f808, plain, (~ (e3 = op(e2, e5)) | ~ sP83), inference(cnf_transformation, [], [f215])).
fof(f215, plain, ((~ (e3 = op(e2, e6)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP83), inference(nnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e3 = op(e2, e6)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP83), inference(usedef, [], [e89])).
fof(e89, plain, (sP83 <=> (~ (e3 = op(e2, e6)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f2048, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2047, f1649])).
fof(f1649, plain, ~ sP84, inference(subsumption_resolution, [], [f801, f357])).
fof(f801, plain, (~ (e3 = op(e5, e2)) | ~ sP84), inference(cnf_transformation, [], [f214])).
fof(f214, plain, ((~ (e3 = op(e6, e2)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP84), inference(nnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e3 = op(e6, e2)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP84), inference(usedef, [], [e90])).
fof(e90, plain, (sP84 <=> (~ (e3 = op(e6, e2)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f2047, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2046, f1647])).
fof(f1647, plain, ~ sP85, inference(subsumption_resolution, [], [f792, f337])).
fof(f792, plain, (~ (e4 = op(e2, e3)) | ~ sP85), inference(cnf_transformation, [], [f213])).
fof(f213, plain, ((~ (e4 = op(e2, e6)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP85), inference(nnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e4 = op(e2, e6)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP85), inference(usedef, [], [e91])).
fof(e91, plain, (sP85 <=> (~ (e4 = op(e2, e6)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f2046, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2045, f1643])).
fof(f1643, plain, ~ sP86, inference(subsumption_resolution, [], [f785, f343])).
fof(f785, plain, (~ (e4 = op(e3, e2)) | ~ sP86), inference(cnf_transformation, [], [f212])).
fof(f212, plain, ((~ (e4 = op(e6, e2)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP86), inference(nnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e4 = op(e6, e2)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP86), inference(usedef, [], [e92])).
fof(e92, plain, (sP86 <=> (~ (e4 = op(e6, e2)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f2045, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2044, f1639])).
fof(f1639, plain, ~ sP87, inference(subsumption_resolution, [], [f779, f338])).
fof(f779, plain, (~ (e5 = op(e2, e4)) | ~ sP87), inference(cnf_transformation, [], [f211])).
fof(f211, plain, ((~ (e5 = op(e2, e6)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e2, e0))) | ~ sP87), inference(nnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e5 = op(e2, e6)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e2, e0))) | ~ sP87), inference(usedef, [], [e93])).
fof(e93, plain, (sP87 <=> (~ (e5 = op(e2, e6)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f2044, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2043, f1636])).
fof(f1636, plain, ~ sP88, inference(subsumption_resolution, [], [f772, f350])).
fof(f772, plain, (~ (e5 = op(e4, e2)) | ~ sP88), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ((~ (e5 = op(e6, e2)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e0, e2))) | ~ sP88), inference(nnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e5 = op(e6, e2)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e0, e2))) | ~ sP88), inference(usedef, [], [e94])).
fof(e94, plain, (sP88 <=> (~ (e5 = op(e6, e2)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f2043, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2042, f1633])).
fof(f1633, plain, ~ sP89, inference(subsumption_resolution, [], [f762, f335])).
fof(f762, plain, (~ (e6 = op(e2, e1)) | ~ sP89), inference(cnf_transformation, [], [f209])).
fof(f209, plain, ((~ (e6 = op(e2, e6)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e2, e0))) | ~ sP89), inference(nnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e6 = op(e2, e6)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e2, e0))) | ~ sP89), inference(usedef, [], [e95])).
fof(e95, plain, (sP89 <=> (~ (e6 = op(e2, e6)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f2042, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2041, f1627])).
fof(f1627, plain, ~ sP90, inference(subsumption_resolution, [], [f755, f329])).
fof(f755, plain, (~ (e6 = op(e1, e2)) | ~ sP90), inference(cnf_transformation, [], [f208])).
fof(f208, plain, ((~ (e6 = op(e6, e2)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e0, e2))) | ~ sP90), inference(nnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e6 = op(e6, e2)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e0, e2))) | ~ sP90), inference(usedef, [], [e96])).
fof(e96, plain, (sP90 <=> (~ (e6 = op(e6, e2)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f2041, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2040, f1621])).
fof(f1621, plain, ~ sP91, inference(subsumption_resolution, [], [f751, f345])).
fof(f751, plain, (~ (e0 = op(e3, e4)) | ~ sP91), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ((~ (e0 = op(e3, e6)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP91), inference(nnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e0 = op(e3, e6)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP91), inference(usedef, [], [e97])).
fof(e97, plain, (sP91 <=> (~ (e0 = op(e3, e6)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f2040, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2039, f1618])).
fof(f1618, plain, ~ sP92, inference(subsumption_resolution, [], [f744, f351])).
fof(f744, plain, (~ (e0 = op(e4, e3)) | ~ sP92), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ((~ (e0 = op(e6, e3)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP92), inference(nnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e0 = op(e6, e3)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP92), inference(usedef, [], [e98])).
fof(e98, plain, (sP92 <=> (~ (e0 = op(e6, e3)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f2039, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2038, f1615])).
fof(f1615, plain, ~ sP93, inference(subsumption_resolution, [], [f739, f347])).
fof(f739, plain, (~ (e1 = op(e3, e6)) | ~ sP93), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ((~ (e1 = op(e3, e6)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP93), inference(nnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e1 = op(e3, e6)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP93), inference(usedef, [], [e99])).
fof(e99, plain, (sP93 <=> (~ (e1 = op(e3, e6)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f2038, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2037, f1614])).
fof(f1614, plain, ~ sP94, inference(subsumption_resolution, [], [f732, f365])).
fof(f732, plain, (~ (e1 = op(e6, e3)) | ~ sP94), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ((~ (e1 = op(e6, e3)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP94), inference(nnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e1 = op(e6, e3)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP94), inference(usedef, [], [e100])).
fof(e100, plain, (sP94 <=> (~ (e1 = op(e6, e3)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f2037, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2036, f1613])).
fof(f1613, plain, ~ sP95, inference(subsumption_resolution, [], [f719, f341])).
fof(f719, plain, (~ (e2 = op(e3, e0)) | ~ sP95), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ((~ (e2 = op(e3, e6)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP95), inference(nnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e2 = op(e3, e6)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP95), inference(usedef, [], [e101])).
fof(e101, plain, (sP95 <=> (~ (e2 = op(e3, e6)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f2036, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2035, f1606])).
fof(f1606, plain, ~ sP96, inference(subsumption_resolution, [], [f712, f323])).
fof(f712, plain, (~ (e2 = op(e0, e3)) | ~ sP96), inference(cnf_transformation, [], [f202])).
fof(f202, plain, ((~ (e2 = op(e6, e3)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP96), inference(nnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e2 = op(e6, e3)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP96), inference(usedef, [], [e102])).
fof(e102, plain, (sP96 <=> (~ (e2 = op(e6, e3)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f2035, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2034, f1599])).
fof(f1599, plain, ~ sP97, inference(subsumption_resolution, [], [f708, f344])).
fof(f708, plain, (~ (e3 = op(e3, e3)) | ~ sP97), inference(cnf_transformation, [], [f201])).
fof(f201, plain, ((~ (e3 = op(e3, e6)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP97), inference(nnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e3 = op(e3, e6)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP97), inference(usedef, [], [e103])).
fof(e103, plain, (sP97 <=> (~ (e3 = op(e3, e6)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f2034, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2033, f1595])).
fof(f1595, plain, ~ sP98, inference(subsumption_resolution, [], [f701, f344])).
fof(f701, plain, (~ (e3 = op(e3, e3)) | ~ sP98), inference(cnf_transformation, [], [f200])).
fof(f200, plain, ((~ (e3 = op(e6, e3)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP98), inference(nnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e3 = op(e6, e3)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP98), inference(usedef, [], [e104])).
fof(e104, plain, (sP98 <=> (~ (e3 = op(e6, e3)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f2033, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2032, f1591])).
fof(f1591, plain, ~ sP99, inference(subsumption_resolution, [], [f693, f343])).
fof(f693, plain, (~ (e4 = op(e3, e2)) | ~ sP99), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ((~ (e4 = op(e3, e6)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP99), inference(nnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e4 = op(e3, e6)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP99), inference(usedef, [], [e105])).
fof(e105, plain, (sP99 <=> (~ (e4 = op(e3, e6)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f2032, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2031, f1586])).
fof(f1586, plain, ~ sP100, inference(subsumption_resolution, [], [f686, f337])).
fof(f686, plain, (~ (e4 = op(e2, e3)) | ~ sP100), inference(cnf_transformation, [], [f198])).
fof(f198, plain, ((~ (e4 = op(e6, e3)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP100), inference(nnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e4 = op(e6, e3)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP100), inference(usedef, [], [e106])).
fof(e106, plain, (sP100 <=> (~ (e4 = op(e6, e3)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f2031, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2030, f1581])).
fof(f1581, plain, ~ sP101, inference(subsumption_resolution, [], [f678, f342])).
fof(f678, plain, (~ (e5 = op(e3, e1)) | ~ sP101), inference(cnf_transformation, [], [f197])).
fof(f197, plain, ((~ (e5 = op(e3, e6)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e3, e0))) | ~ sP101), inference(nnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e5 = op(e3, e6)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e3, e0))) | ~ sP101), inference(usedef, [], [e107])).
fof(e107, plain, (sP101 <=> (~ (e5 = op(e3, e6)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f2030, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2029, f1575])).
fof(f1575, plain, ~ sP102, inference(subsumption_resolution, [], [f671, f330])).
fof(f671, plain, (~ (e5 = op(e1, e3)) | ~ sP102), inference(cnf_transformation, [], [f196])).
fof(f196, plain, ((~ (e5 = op(e6, e3)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e0, e3))) | ~ sP102), inference(nnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e5 = op(e6, e3)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e0, e3))) | ~ sP102), inference(usedef, [], [e108])).
fof(e108, plain, (sP102 <=> (~ (e5 = op(e6, e3)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f2029, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2028, f1569])).
fof(f1569, plain, ~ sP103, inference(subsumption_resolution, [], [f668, f346])).
fof(f668, plain, (~ (e6 = op(e3, e5)) | ~ sP103), inference(cnf_transformation, [], [f195])).
fof(f195, plain, ((~ (e6 = op(e3, e6)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e3, e0))) | ~ sP103), inference(nnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e6 = op(e3, e6)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e3, e0))) | ~ sP103), inference(usedef, [], [e109])).
fof(e109, plain, (sP103 <=> (~ (e6 = op(e3, e6)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f2028, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2027, f1567])).
fof(f1567, plain, ~ sP104, inference(subsumption_resolution, [], [f661, f358])).
fof(f661, plain, (~ (e6 = op(e5, e3)) | ~ sP104), inference(cnf_transformation, [], [f194])).
fof(f194, plain, ((~ (e6 = op(e6, e3)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e0, e3))) | ~ sP104), inference(nnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e6 = op(e6, e3)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e0, e3))) | ~ sP104), inference(usedef, [], [e110])).
fof(e110, plain, (sP104 <=> (~ (e6 = op(e6, e3)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f2027, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2026, f1565])).
fof(f1565, plain, ~ sP105, inference(subsumption_resolution, [], [f652, f351])).
fof(f652, plain, (~ (e0 = op(e4, e3)) | ~ sP105), inference(cnf_transformation, [], [f193])).
fof(f193, plain, ((~ (e0 = op(e4, e6)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP105), inference(nnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e0 = op(e4, e6)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP105), inference(usedef, [], [e111])).
fof(e111, plain, (sP105 <=> (~ (e0 = op(e4, e6)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f2026, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2025, f1561])).
fof(f1561, plain, ~ sP106, inference(subsumption_resolution, [], [f645, f345])).
fof(f645, plain, (~ (e0 = op(e3, e4)) | ~ sP106), inference(cnf_transformation, [], [f192])).
fof(f192, plain, ((~ (e0 = op(e6, e4)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP106), inference(nnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e0 = op(e6, e4)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP106), inference(usedef, [], [e112])).
fof(e112, plain, (sP106 <=> (~ (e0 = op(e6, e4)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f2025, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2024, f1557])).
fof(f1557, plain, ~ sP107, inference(subsumption_resolution, [], [f640, f353])).
fof(f640, plain, (~ (e1 = op(e4, e5)) | ~ sP107), inference(cnf_transformation, [], [f191])).
fof(f191, plain, ((~ (e1 = op(e4, e6)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP107), inference(nnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e1 = op(e4, e6)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP107), inference(usedef, [], [e113])).
fof(e113, plain, (sP107 <=> (~ (e1 = op(e4, e6)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f2024, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2023, f1555])).
fof(f1555, plain, ~ sP108, inference(subsumption_resolution, [], [f633, f359])).
fof(f633, plain, (~ (e1 = op(e5, e4)) | ~ sP108), inference(cnf_transformation, [], [f190])).
fof(f190, plain, ((~ (e1 = op(e6, e4)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP108), inference(nnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e1 = op(e6, e4)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP108), inference(usedef, [], [e114])).
fof(e114, plain, (sP108 <=> (~ (e1 = op(e6, e4)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f2023, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2022, f1553])).
fof(f1553, plain, ~ sP109, inference(subsumption_resolution, [], [f622, f349])).
fof(f622, plain, (~ (e2 = op(e4, e1)) | ~ sP109), inference(cnf_transformation, [], [f189])).
fof(f189, plain, ((~ (e2 = op(e4, e6)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP109), inference(nnf_transformation, [], [f115])).
fof(f115, plain, ((~ (e2 = op(e4, e6)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP109), inference(usedef, [], [e115])).
fof(e115, plain, (sP109 <=> (~ (e2 = op(e4, e6)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f2022, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2021, f1547])).
fof(f1547, plain, ~ sP110, inference(subsumption_resolution, [], [f615, f331])).
fof(f615, plain, (~ (e2 = op(e1, e4)) | ~ sP110), inference(cnf_transformation, [], [f188])).
fof(f188, plain, ((~ (e2 = op(e6, e4)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP110), inference(nnf_transformation, [], [f116])).
fof(f116, plain, ((~ (e2 = op(e6, e4)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP110), inference(usedef, [], [e116])).
fof(e116, plain, (sP110 <=> (~ (e2 = op(e6, e4)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f2021, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2020, f1541])).
fof(f1541, plain, ~ sP111, inference(subsumption_resolution, [], [f613, f354])).
fof(f613, plain, (~ (e3 = op(e4, e6)) | ~ sP111), inference(cnf_transformation, [], [f187])).
fof(f187, plain, ((~ (e3 = op(e4, e6)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP111), inference(nnf_transformation, [], [f117])).
fof(f117, plain, ((~ (e3 = op(e4, e6)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP111), inference(usedef, [], [e117])).
fof(e117, plain, (sP111 <=> (~ (e3 = op(e4, e6)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f2020, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2019, f1540])).
fof(f1540, plain, ~ sP112, inference(subsumption_resolution, [], [f606, f366])).
fof(f606, plain, (~ (e3 = op(e6, e4)) | ~ sP112), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e3 = op(e6, e4)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP112), inference(nnf_transformation, [], [f118])).
fof(f118, plain, ((~ (e3 = op(e6, e4)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP112), inference(usedef, [], [e118])).
fof(e118, plain, (sP112 <=> (~ (e3 = op(e6, e4)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f2019, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2018, f1539])).
fof(f1539, plain, ~ sP113, inference(subsumption_resolution, [], [f597, f352])).
fof(f597, plain, (~ (e4 = op(e4, e4)) | ~ sP113), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e4 = op(e4, e6)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP113), inference(nnf_transformation, [], [f119])).
fof(f119, plain, ((~ (e4 = op(e4, e6)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP113), inference(usedef, [], [e119])).
fof(e119, plain, (sP113 <=> (~ (e4 = op(e4, e6)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f2018, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2017, f1536])).
fof(f1536, plain, ~ sP114, inference(subsumption_resolution, [], [f590, f352])).
fof(f590, plain, (~ (e4 = op(e4, e4)) | ~ sP114), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e4 = op(e6, e4)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | ~ sP114), inference(nnf_transformation, [], [f120])).
fof(f120, plain, ((~ (e4 = op(e6, e4)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | ~ sP114), inference(usedef, [], [e120])).
fof(e120, plain, (sP114 <=> (~ (e4 = op(e6, e4)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f2017, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2016, f1533])).
fof(f1533, plain, ~ sP115, inference(subsumption_resolution, [], [f581, f350])).
fof(f581, plain, (~ (e5 = op(e4, e2)) | ~ sP115), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ((~ (e5 = op(e4, e6)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e4, e0))) | ~ sP115), inference(nnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e5 = op(e4, e6)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e4, e0))) | ~ sP115), inference(usedef, [], [e121])).
fof(e121, plain, (sP115 <=> (~ (e5 = op(e4, e6)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f2016, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2015, f1528])).
fof(f1528, plain, ~ sP116, inference(subsumption_resolution, [], [f574, f338])).
fof(f574, plain, (~ (e5 = op(e2, e4)) | ~ sP116), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((~ (e5 = op(e6, e4)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e0, e4))) | ~ sP116), inference(nnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e5 = op(e6, e4)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e0, e4))) | ~ sP116), inference(usedef, [], [e122])).
fof(e122, plain, (sP116 <=> (~ (e5 = op(e6, e4)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f2015, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2014, f1523])).
fof(f1523, plain, ~ sP117, inference(subsumption_resolution, [], [f565, f348])).
fof(f565, plain, (~ (e6 = op(e4, e0)) | ~ sP117), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ((~ (e6 = op(e4, e6)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e4, e0))) | ~ sP117), inference(nnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e6 = op(e4, e6)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e4, e0))) | ~ sP117), inference(usedef, [], [e123])).
fof(e123, plain, (sP117 <=> (~ (e6 = op(e4, e6)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f2014, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2013, f1516])).
fof(f1516, plain, ~ sP118, inference(subsumption_resolution, [], [f558, f324])).
fof(f558, plain, (~ (e6 = op(e0, e4)) | ~ sP118), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((~ (e6 = op(e6, e4)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e0, e4))) | ~ sP118), inference(nnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e6 = op(e6, e4)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e0, e4))) | ~ sP118), inference(usedef, [], [e124])).
fof(e124, plain, (sP118 <=> (~ (e6 = op(e6, e4)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f2013, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2012, f1509])).
fof(f1509, plain, ~ sP119, inference(subsumption_resolution, [], [f552, f356])).
fof(f552, plain, (~ (e0 = op(e5, e1)) | ~ sP119), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((~ (e0 = op(e5, e6)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e5, e0))) | ~ sP119), inference(nnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e0 = op(e5, e6)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e5, e0))) | ~ sP119), inference(usedef, [], [e125])).
fof(e125, plain, (sP119 <=> (~ (e0 = op(e5, e6)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f2012, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2011, f1503])).
fof(f1503, plain, ~ sP120, inference(subsumption_resolution, [], [f545, f332])).
fof(f545, plain, (~ (e0 = op(e1, e5)) | ~ sP120), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ((~ (e0 = op(e6, e5)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e0, e5))) | ~ sP120), inference(nnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e0 = op(e6, e5)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e0, e5))) | ~ sP120), inference(usedef, [], [e126])).
fof(e126, plain, (sP120 <=> (~ (e0 = op(e6, e5)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f2011, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2010, f1497])).
fof(f1497, plain, ~ sP121, inference(subsumption_resolution, [], [f541, f359])).
fof(f541, plain, (~ (e1 = op(e5, e4)) | ~ sP121), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ((~ (e1 = op(e5, e6)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e5, e0))) | ~ sP121), inference(nnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e1 = op(e5, e6)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e5, e0))) | ~ sP121), inference(usedef, [], [e127])).
fof(e127, plain, (sP121 <=> (~ (e1 = op(e5, e6)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f2010, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2009, f1494])).
fof(f1494, plain, ~ sP122, inference(subsumption_resolution, [], [f534, f353])).
fof(f534, plain, (~ (e1 = op(e4, e5)) | ~ sP122), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ((~ (e1 = op(e6, e5)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e0, e5))) | ~ sP122), inference(nnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e1 = op(e6, e5)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e0, e5))) | ~ sP122), inference(usedef, [], [e128])).
fof(e128, plain, (sP122 <=> (~ (e1 = op(e6, e5)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f2009, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2008, f1491])).
fof(f1491, plain, ~ sP123, inference(subsumption_resolution, [], [f529, f361])).
fof(f529, plain, (~ (e2 = op(e5, e6)) | ~ sP123), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e2 = op(e5, e6)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e5, e0))) | ~ sP123), inference(nnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e2 = op(e5, e6)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e5, e0))) | ~ sP123), inference(usedef, [], [e129])).
fof(e129, plain, (sP123 <=> (~ (e2 = op(e5, e6)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f2008, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2007, f1490])).
fof(f1490, plain, ~ sP124, inference(subsumption_resolution, [], [f522, f367])).
fof(f522, plain, (~ (e2 = op(e6, e5)) | ~ sP124), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e2 = op(e6, e5)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e0, e5))) | ~ sP124), inference(nnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e2 = op(e6, e5)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e0, e5))) | ~ sP124), inference(usedef, [], [e130])).
fof(e130, plain, (sP124 <=> (~ (e2 = op(e6, e5)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f2007, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2006, f1489])).
fof(f1489, plain, ~ sP125, inference(subsumption_resolution, [], [f511, f357])).
fof(f511, plain, (~ (e3 = op(e5, e2)) | ~ sP125), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e3 = op(e5, e6)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e5, e0))) | ~ sP125), inference(nnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e3 = op(e5, e6)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e5, e0))) | ~ sP125), inference(usedef, [], [e131])).
fof(e131, plain, (sP125 <=> (~ (e3 = op(e5, e6)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f2006, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2005, f1484])).
fof(f1484, plain, ~ sP126, inference(subsumption_resolution, [], [f504, f339])).
fof(f504, plain, (~ (e3 = op(e2, e5)) | ~ sP126), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e3 = op(e6, e5)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e0, e5))) | ~ sP126), inference(nnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e3 = op(e6, e5)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e0, e5))) | ~ sP126), inference(usedef, [], [e132])).
fof(e132, plain, (sP126 <=> (~ (e3 = op(e6, e5)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f2005, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2004, f1479])).
fof(f1479, plain, ~ sP127, inference(subsumption_resolution, [], [f495, f355])).
fof(f495, plain, (~ (e4 = op(e5, e0)) | ~ sP127), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e4 = op(e5, e6)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e5, e0))) | ~ sP127), inference(nnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e4 = op(e5, e6)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e5, e0))) | ~ sP127), inference(usedef, [], [e133])).
fof(e133, plain, (sP127 <=> (~ (e4 = op(e5, e6)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f2004, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2003, f1472])).
fof(f1472, plain, ~ sP128, inference(subsumption_resolution, [], [f488, f325])).
fof(f488, plain, (~ (e4 = op(e0, e5)) | ~ sP128), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((~ (e4 = op(e6, e5)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e0, e5))) | ~ sP128), inference(nnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e4 = op(e6, e5)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e0, e5))) | ~ sP128), inference(usedef, [], [e134])).
fof(e134, plain, (sP128 <=> (~ (e4 = op(e6, e5)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f2003, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2002, f1465])).
fof(f1465, plain, ~ sP129, inference(subsumption_resolution, [], [f486, f360])).
fof(f486, plain, (~ (e5 = op(e5, e5)) | ~ sP129), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((~ (e5 = op(e5, e6)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e5, e0))) | ~ sP129), inference(nnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e5 = op(e5, e6)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e5, e0))) | ~ sP129), inference(usedef, [], [e135])).
fof(e135, plain, (sP129 <=> (~ (e5 = op(e5, e6)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f2002, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2001, f1463])).
fof(f1463, plain, ~ sP130, inference(subsumption_resolution, [], [f479, f360])).
fof(f479, plain, (~ (e5 = op(e5, e5)) | ~ sP130), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((~ (e5 = op(e6, e5)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e0, e5))) | ~ sP130), inference(nnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e5 = op(e6, e5)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e0, e5))) | ~ sP130), inference(usedef, [], [e136])).
fof(e136, plain, (sP130 <=> (~ (e5 = op(e6, e5)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP130])])).
fof(f2001, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f2000, f1461])).
fof(f1461, plain, ~ sP131, inference(subsumption_resolution, [], [f470, f358])).
fof(f470, plain, (~ (e6 = op(e5, e3)) | ~ sP131), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e6 = op(e5, e6)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e5, e0))) | ~ sP131), inference(nnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e6 = op(e5, e6)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e5, e0))) | ~ sP131), inference(usedef, [], [e137])).
fof(e137, plain, (sP131 <=> (~ (e6 = op(e5, e6)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP131])])).
fof(f2000, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1999, f1457])).
fof(f1457, plain, ~ sP132, inference(subsumption_resolution, [], [f463, f346])).
fof(f463, plain, (~ (e6 = op(e3, e5)) | ~ sP132), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e6 = op(e6, e5)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e0, e5))) | ~ sP132), inference(nnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e6 = op(e6, e5)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e0, e5))) | ~ sP132), inference(usedef, [], [e138])).
fof(e138, plain, (sP132 <=> (~ (e6 = op(e6, e5)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP132])])).
fof(f1999, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1998, f1453])).
fof(f1453, plain, ~ sP133, inference(subsumption_resolution, [], [f455, f364])).
fof(f455, plain, (~ (e0 = op(e6, e2)) | ~ sP133), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e0 = op(e6, e6)) & ~ (e0 = op(e6, e5)) & ~ (e0 = op(e6, e4)) & ~ (e0 = op(e6, e3)) & ~ (e0 = op(e6, e2)) & ~ (e0 = op(e6, e1)) & ~ (e0 = op(e6, e0))) | ~ sP133), inference(nnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e0 = op(e6, e6)) & ~ (e0 = op(e6, e5)) & ~ (e0 = op(e6, e4)) & ~ (e0 = op(e6, e3)) & ~ (e0 = op(e6, e2)) & ~ (e0 = op(e6, e1)) & ~ (e0 = op(e6, e0))) | ~ sP133), inference(usedef, [], [e139])).
fof(e139, plain, (sP133 <=> (~ (e0 = op(e6, e6)) & ~ (e0 = op(e6, e5)) & ~ (e0 = op(e6, e4)) & ~ (e0 = op(e6, e3)) & ~ (e0 = op(e6, e2)) & ~ (e0 = op(e6, e1)) & ~ (e0 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP133])])).
fof(f1998, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1997, f1448])).
fof(f1448, plain, ~ sP134, inference(subsumption_resolution, [], [f448, f340])).
fof(f448, plain, (~ (e0 = op(e2, e6)) | ~ sP134), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e0 = op(e6, e6)) & ~ (e0 = op(e5, e6)) & ~ (e0 = op(e4, e6)) & ~ (e0 = op(e3, e6)) & ~ (e0 = op(e2, e6)) & ~ (e0 = op(e1, e6)) & ~ (e0 = op(e0, e6))) | ~ sP134), inference(nnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e0 = op(e6, e6)) & ~ (e0 = op(e5, e6)) & ~ (e0 = op(e4, e6)) & ~ (e0 = op(e3, e6)) & ~ (e0 = op(e2, e6)) & ~ (e0 = op(e1, e6)) & ~ (e0 = op(e0, e6))) | ~ sP134), inference(usedef, [], [e140])).
fof(e140, plain, (sP134 <=> (~ (e0 = op(e6, e6)) & ~ (e0 = op(e5, e6)) & ~ (e0 = op(e4, e6)) & ~ (e0 = op(e3, e6)) & ~ (e0 = op(e2, e6)) & ~ (e0 = op(e1, e6)) & ~ (e0 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP134])])).
fof(f1997, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1996, f1443])).
fof(f1443, plain, ~ sP135, inference(subsumption_resolution, [], [f442, f365])).
fof(f442, plain, (~ (e1 = op(e6, e3)) | ~ sP135), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e1 = op(e6, e6)) & ~ (e1 = op(e6, e5)) & ~ (e1 = op(e6, e4)) & ~ (e1 = op(e6, e3)) & ~ (e1 = op(e6, e2)) & ~ (e1 = op(e6, e1)) & ~ (e1 = op(e6, e0))) | ~ sP135), inference(nnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e1 = op(e6, e6)) & ~ (e1 = op(e6, e5)) & ~ (e1 = op(e6, e4)) & ~ (e1 = op(e6, e3)) & ~ (e1 = op(e6, e2)) & ~ (e1 = op(e6, e1)) & ~ (e1 = op(e6, e0))) | ~ sP135), inference(usedef, [], [e141])).
fof(e141, plain, (sP135 <=> (~ (e1 = op(e6, e6)) & ~ (e1 = op(e6, e5)) & ~ (e1 = op(e6, e4)) & ~ (e1 = op(e6, e3)) & ~ (e1 = op(e6, e2)) & ~ (e1 = op(e6, e1)) & ~ (e1 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP135])])).
fof(f1996, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1995, f1439])).
fof(f1439, plain, ~ sP136, inference(subsumption_resolution, [], [f435, f347])).
fof(f435, plain, (~ (e1 = op(e3, e6)) | ~ sP136), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e1 = op(e6, e6)) & ~ (e1 = op(e5, e6)) & ~ (e1 = op(e4, e6)) & ~ (e1 = op(e3, e6)) & ~ (e1 = op(e2, e6)) & ~ (e1 = op(e1, e6)) & ~ (e1 = op(e0, e6))) | ~ sP136), inference(nnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e1 = op(e6, e6)) & ~ (e1 = op(e5, e6)) & ~ (e1 = op(e4, e6)) & ~ (e1 = op(e3, e6)) & ~ (e1 = op(e2, e6)) & ~ (e1 = op(e1, e6)) & ~ (e1 = op(e0, e6))) | ~ sP136), inference(usedef, [], [e142])).
fof(e142, plain, (sP136 <=> (~ (e1 = op(e6, e6)) & ~ (e1 = op(e5, e6)) & ~ (e1 = op(e4, e6)) & ~ (e1 = op(e3, e6)) & ~ (e1 = op(e2, e6)) & ~ (e1 = op(e1, e6)) & ~ (e1 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP136])])).
fof(f1995, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1994, f1435])).
fof(f1435, plain, ~ sP137, inference(subsumption_resolution, [], [f430, f367])).
fof(f430, plain, (~ (e2 = op(e6, e5)) | ~ sP137), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e2 = op(e6, e6)) & ~ (e2 = op(e6, e5)) & ~ (e2 = op(e6, e4)) & ~ (e2 = op(e6, e3)) & ~ (e2 = op(e6, e2)) & ~ (e2 = op(e6, e1)) & ~ (e2 = op(e6, e0))) | ~ sP137), inference(nnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e2 = op(e6, e6)) & ~ (e2 = op(e6, e5)) & ~ (e2 = op(e6, e4)) & ~ (e2 = op(e6, e3)) & ~ (e2 = op(e6, e2)) & ~ (e2 = op(e6, e1)) & ~ (e2 = op(e6, e0))) | ~ sP137), inference(usedef, [], [e143])).
fof(e143, plain, (sP137 <=> (~ (e2 = op(e6, e6)) & ~ (e2 = op(e6, e5)) & ~ (e2 = op(e6, e4)) & ~ (e2 = op(e6, e3)) & ~ (e2 = op(e6, e2)) & ~ (e2 = op(e6, e1)) & ~ (e2 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP137])])).
fof(f1994, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1993, f1433])).
fof(f1433, plain, ~ sP138, inference(subsumption_resolution, [], [f423, f361])).
fof(f423, plain, (~ (e2 = op(e5, e6)) | ~ sP138), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e2 = op(e6, e6)) & ~ (e2 = op(e5, e6)) & ~ (e2 = op(e4, e6)) & ~ (e2 = op(e3, e6)) & ~ (e2 = op(e2, e6)) & ~ (e2 = op(e1, e6)) & ~ (e2 = op(e0, e6))) | ~ sP138), inference(nnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e2 = op(e6, e6)) & ~ (e2 = op(e5, e6)) & ~ (e2 = op(e4, e6)) & ~ (e2 = op(e3, e6)) & ~ (e2 = op(e2, e6)) & ~ (e2 = op(e1, e6)) & ~ (e2 = op(e0, e6))) | ~ sP138), inference(usedef, [], [e144])).
fof(e144, plain, (sP138 <=> (~ (e2 = op(e6, e6)) & ~ (e2 = op(e5, e6)) & ~ (e2 = op(e4, e6)) & ~ (e2 = op(e3, e6)) & ~ (e2 = op(e2, e6)) & ~ (e2 = op(e1, e6)) & ~ (e2 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP138])])).
fof(f1993, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1992, f1431])).
fof(f1431, plain, ~ sP139, inference(subsumption_resolution, [], [f415, f366])).
fof(f415, plain, (~ (e3 = op(e6, e4)) | ~ sP139), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e3 = op(e6, e6)) & ~ (e3 = op(e6, e5)) & ~ (e3 = op(e6, e4)) & ~ (e3 = op(e6, e3)) & ~ (e3 = op(e6, e2)) & ~ (e3 = op(e6, e1)) & ~ (e3 = op(e6, e0))) | ~ sP139), inference(nnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e3 = op(e6, e6)) & ~ (e3 = op(e6, e5)) & ~ (e3 = op(e6, e4)) & ~ (e3 = op(e6, e3)) & ~ (e3 = op(e6, e2)) & ~ (e3 = op(e6, e1)) & ~ (e3 = op(e6, e0))) | ~ sP139), inference(usedef, [], [e145])).
fof(e145, plain, (sP139 <=> (~ (e3 = op(e6, e6)) & ~ (e3 = op(e6, e5)) & ~ (e3 = op(e6, e4)) & ~ (e3 = op(e6, e3)) & ~ (e3 = op(e6, e2)) & ~ (e3 = op(e6, e1)) & ~ (e3 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP139])])).
fof(f1992, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1991, f1428])).
fof(f1428, plain, ~ sP140, inference(subsumption_resolution, [], [f408, f354])).
fof(f408, plain, (~ (e3 = op(e4, e6)) | ~ sP140), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e3 = op(e6, e6)) & ~ (e3 = op(e5, e6)) & ~ (e3 = op(e4, e6)) & ~ (e3 = op(e3, e6)) & ~ (e3 = op(e2, e6)) & ~ (e3 = op(e1, e6)) & ~ (e3 = op(e0, e6))) | ~ sP140), inference(nnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e3 = op(e6, e6)) & ~ (e3 = op(e5, e6)) & ~ (e3 = op(e4, e6)) & ~ (e3 = op(e3, e6)) & ~ (e3 = op(e2, e6)) & ~ (e3 = op(e1, e6)) & ~ (e3 = op(e0, e6))) | ~ sP140), inference(usedef, [], [e146])).
fof(e146, plain, (sP140 <=> (~ (e3 = op(e6, e6)) & ~ (e3 = op(e5, e6)) & ~ (e3 = op(e4, e6)) & ~ (e3 = op(e3, e6)) & ~ (e3 = op(e2, e6)) & ~ (e3 = op(e1, e6)) & ~ (e3 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP140])])).
fof(f1991, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1990, f1425])).
fof(f1425, plain, ~ sP141, inference(subsumption_resolution, [], [f398, f363])).
fof(f398, plain, (~ (e4 = op(e6, e1)) | ~ sP141), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e4 = op(e6, e6)) & ~ (e4 = op(e6, e5)) & ~ (e4 = op(e6, e4)) & ~ (e4 = op(e6, e3)) & ~ (e4 = op(e6, e2)) & ~ (e4 = op(e6, e1)) & ~ (e4 = op(e6, e0))) | ~ sP141), inference(nnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e4 = op(e6, e6)) & ~ (e4 = op(e6, e5)) & ~ (e4 = op(e6, e4)) & ~ (e4 = op(e6, e3)) & ~ (e4 = op(e6, e2)) & ~ (e4 = op(e6, e1)) & ~ (e4 = op(e6, e0))) | ~ sP141), inference(usedef, [], [e147])).
fof(e147, plain, (sP141 <=> (~ (e4 = op(e6, e6)) & ~ (e4 = op(e6, e5)) & ~ (e4 = op(e6, e4)) & ~ (e4 = op(e6, e3)) & ~ (e4 = op(e6, e2)) & ~ (e4 = op(e6, e1)) & ~ (e4 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP141])])).
fof(f1990, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1989, f1419])).
fof(f1419, plain, ~ sP142, inference(subsumption_resolution, [], [f391, f333])).
fof(f391, plain, (~ (e4 = op(e1, e6)) | ~ sP142), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e4 = op(e6, e6)) & ~ (e4 = op(e5, e6)) & ~ (e4 = op(e4, e6)) & ~ (e4 = op(e3, e6)) & ~ (e4 = op(e2, e6)) & ~ (e4 = op(e1, e6)) & ~ (e4 = op(e0, e6))) | ~ sP142), inference(nnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e4 = op(e6, e6)) & ~ (e4 = op(e5, e6)) & ~ (e4 = op(e4, e6)) & ~ (e4 = op(e3, e6)) & ~ (e4 = op(e2, e6)) & ~ (e4 = op(e1, e6)) & ~ (e4 = op(e0, e6))) | ~ sP142), inference(usedef, [], [e148])).
fof(e148, plain, (sP142 <=> (~ (e4 = op(e6, e6)) & ~ (e4 = op(e5, e6)) & ~ (e4 = op(e4, e6)) & ~ (e4 = op(e3, e6)) & ~ (e4 = op(e2, e6)) & ~ (e4 = op(e1, e6)) & ~ (e4 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP142])])).
fof(f1989, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1988, f1413])).
fof(f1413, plain, ~ sP143, inference(subsumption_resolution, [], [f383, f362])).
fof(f383, plain, (~ (e5 = op(e6, e0)) | ~ sP143), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e5 = op(e6, e6)) & ~ (e5 = op(e6, e5)) & ~ (e5 = op(e6, e4)) & ~ (e5 = op(e6, e3)) & ~ (e5 = op(e6, e2)) & ~ (e5 = op(e6, e1)) & ~ (e5 = op(e6, e0))) | ~ sP143), inference(nnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e5 = op(e6, e6)) & ~ (e5 = op(e6, e5)) & ~ (e5 = op(e6, e4)) & ~ (e5 = op(e6, e3)) & ~ (e5 = op(e6, e2)) & ~ (e5 = op(e6, e1)) & ~ (e5 = op(e6, e0))) | ~ sP143), inference(usedef, [], [e149])).
fof(e149, plain, (sP143 <=> (~ (e5 = op(e6, e6)) & ~ (e5 = op(e6, e5)) & ~ (e5 = op(e6, e4)) & ~ (e5 = op(e6, e3)) & ~ (e5 = op(e6, e2)) & ~ (e5 = op(e6, e1)) & ~ (e5 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP143])])).
fof(f1988, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1987, f1406])).
fof(f1406, plain, ~ sP144, inference(subsumption_resolution, [], [f376, f326])).
fof(f376, plain, (~ (e5 = op(e0, e6)) | ~ sP144), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ((~ (e5 = op(e6, e6)) & ~ (e5 = op(e5, e6)) & ~ (e5 = op(e4, e6)) & ~ (e5 = op(e3, e6)) & ~ (e5 = op(e2, e6)) & ~ (e5 = op(e1, e6)) & ~ (e5 = op(e0, e6))) | ~ sP144), inference(nnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e5 = op(e6, e6)) & ~ (e5 = op(e5, e6)) & ~ (e5 = op(e4, e6)) & ~ (e5 = op(e3, e6)) & ~ (e5 = op(e2, e6)) & ~ (e5 = op(e1, e6)) & ~ (e5 = op(e0, e6))) | ~ sP144), inference(usedef, [], [e150])).
fof(e150, plain, (sP144 <=> (~ (e5 = op(e6, e6)) & ~ (e5 = op(e5, e6)) & ~ (e5 = op(e4, e6)) & ~ (e5 = op(e3, e6)) & ~ (e5 = op(e2, e6)) & ~ (e5 = op(e1, e6)) & ~ (e5 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP144])])).
fof(f1987, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1986, f1399])).
fof(f1399, plain, ~ sP145, inference(subsumption_resolution, [], [f375, f368])).
fof(f375, plain, (~ (e6 = op(e6, e6)) | ~ sP145), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e6 = op(e6, e6)) & ~ (e6 = op(e6, e5)) & ~ (e6 = op(e6, e4)) & ~ (e6 = op(e6, e3)) & ~ (e6 = op(e6, e2)) & ~ (e6 = op(e6, e1)) & ~ (e6 = op(e6, e0))) | ~ sP145), inference(nnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e6 = op(e6, e6)) & ~ (e6 = op(e6, e5)) & ~ (e6 = op(e6, e4)) & ~ (e6 = op(e6, e3)) & ~ (e6 = op(e6, e2)) & ~ (e6 = op(e6, e1)) & ~ (e6 = op(e6, e0))) | ~ sP145), inference(usedef, [], [e151])).
fof(e151, plain, (sP145 <=> (~ (e6 = op(e6, e6)) & ~ (e6 = op(e6, e5)) & ~ (e6 = op(e6, e4)) & ~ (e6 = op(e6, e3)) & ~ (e6 = op(e6, e2)) & ~ (e6 = op(e6, e1)) & ~ (e6 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP145])])).
fof(f1986, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP145 | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(subsumption_resolution, [], [f1398, f368])).
fof(f1398, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | ~ (e6 = op(e6, e6)) | sP145 | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(duplicate_literal_removal, [], [f1397])).
fof(f1397, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | ~ (e6 = op(e6, e6)) | sP145 | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e6 = op(e6, e6)) | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(cnf_transformation, [], [f152])).
fof(f152, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | (~ (e6 = op(e6, e6)) & ~ (e6 = op(e5, e6)) & ~ (e6 = op(e4, e6)) & ~ (e6 = op(e3, e6)) & ~ (e6 = op(e2, e6)) & ~ (e6 = op(e1, e6)) & ~ (e6 = op(e0, e6))) | sP145 | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0 | ~ (e6 = op(e6, e6)) | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(definition_folding, [], [f5, e151, e150, e149, e148, e147, e146, e145, e144, e143, e142, e141, e140, e139, e138, e137, e136, e135, e134, e133, e132, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7, e6])).
fof(f5, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | (~ (e6 = op(e6, e6)) & ~ (e6 = op(e5, e6)) & ~ (e6 = op(e4, e6)) & ~ (e6 = op(e3, e6)) & ~ (e6 = op(e2, e6)) & ~ (e6 = op(e1, e6)) & ~ (e6 = op(e0, e6))) | (~ (e6 = op(e6, e6)) & ~ (e6 = op(e6, e5)) & ~ (e6 = op(e6, e4)) & ~ (e6 = op(e6, e3)) & ~ (e6 = op(e6, e2)) & ~ (e6 = op(e6, e1)) & ~ (e6 = op(e6, e0))) | (~ (e5 = op(e6, e6)) & ~ (e5 = op(e5, e6)) & ~ (e5 = op(e4, e6)) & ~ (e5 = op(e3, e6)) & ~ (e5 = op(e2, e6)) & ~ (e5 = op(e1, e6)) & ~ (e5 = op(e0, e6))) | (~ (e5 = op(e6, e6)) & ~ (e5 = op(e6, e5)) & ~ (e5 = op(e6, e4)) & ~ (e5 = op(e6, e3)) & ~ (e5 = op(e6, e2)) & ~ (e5 = op(e6, e1)) & ~ (e5 = op(e6, e0))) | (~ (e4 = op(e6, e6)) & ~ (e4 = op(e5, e6)) & ~ (e4 = op(e4, e6)) & ~ (e4 = op(e3, e6)) & ~ (e4 = op(e2, e6)) & ~ (e4 = op(e1, e6)) & ~ (e4 = op(e0, e6))) | (~ (e4 = op(e6, e6)) & ~ (e4 = op(e6, e5)) & ~ (e4 = op(e6, e4)) & ~ (e4 = op(e6, e3)) & ~ (e4 = op(e6, e2)) & ~ (e4 = op(e6, e1)) & ~ (e4 = op(e6, e0))) | (~ (e3 = op(e6, e6)) & ~ (e3 = op(e5, e6)) & ~ (e3 = op(e4, e6)) & ~ (e3 = op(e3, e6)) & ~ (e3 = op(e2, e6)) & ~ (e3 = op(e1, e6)) & ~ (e3 = op(e0, e6))) | (~ (e3 = op(e6, e6)) & ~ (e3 = op(e6, e5)) & ~ (e3 = op(e6, e4)) & ~ (e3 = op(e6, e3)) & ~ (e3 = op(e6, e2)) & ~ (e3 = op(e6, e1)) & ~ (e3 = op(e6, e0))) | (~ (e2 = op(e6, e6)) & ~ (e2 = op(e5, e6)) & ~ (e2 = op(e4, e6)) & ~ (e2 = op(e3, e6)) & ~ (e2 = op(e2, e6)) & ~ (e2 = op(e1, e6)) & ~ (e2 = op(e0, e6))) | (~ (e2 = op(e6, e6)) & ~ (e2 = op(e6, e5)) & ~ (e2 = op(e6, e4)) & ~ (e2 = op(e6, e3)) & ~ (e2 = op(e6, e2)) & ~ (e2 = op(e6, e1)) & ~ (e2 = op(e6, e0))) | (~ (e1 = op(e6, e6)) & ~ (e1 = op(e5, e6)) & ~ (e1 = op(e4, e6)) & ~ (e1 = op(e3, e6)) & ~ (e1 = op(e2, e6)) & ~ (e1 = op(e1, e6)) & ~ (e1 = op(e0, e6))) | (~ (e1 = op(e6, e6)) & ~ (e1 = op(e6, e5)) & ~ (e1 = op(e6, e4)) & ~ (e1 = op(e6, e3)) & ~ (e1 = op(e6, e2)) & ~ (e1 = op(e6, e1)) & ~ (e1 = op(e6, e0))) | (~ (e0 = op(e6, e6)) & ~ (e0 = op(e5, e6)) & ~ (e0 = op(e4, e6)) & ~ (e0 = op(e3, e6)) & ~ (e0 = op(e2, e6)) & ~ (e0 = op(e1, e6)) & ~ (e0 = op(e0, e6))) | (~ (e0 = op(e6, e6)) & ~ (e0 = op(e6, e5)) & ~ (e0 = op(e6, e4)) & ~ (e0 = op(e6, e3)) & ~ (e0 = op(e6, e2)) & ~ (e0 = op(e6, e1)) & ~ (e0 = op(e6, e0))) | (~ (e6 = op(e6, e5)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e0, e5))) | (~ (e6 = op(e5, e6)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e5, e0))) | (~ (e5 = op(e6, e5)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e0, e5))) | (~ (e5 = op(e5, e6)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e5, e0))) | (~ (e4 = op(e6, e5)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e0, e5))) | (~ (e4 = op(e5, e6)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e5, e0))) | (~ (e3 = op(e6, e5)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e0, e5))) | (~ (e3 = op(e5, e6)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e5, e0))) | (~ (e2 = op(e6, e5)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e0, e5))) | (~ (e2 = op(e5, e6)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e5, e0))) | (~ (e1 = op(e6, e5)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e0, e5))) | (~ (e1 = op(e5, e6)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e5, e0))) | (~ (e0 = op(e6, e5)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e0, e5))) | (~ (e0 = op(e5, e6)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e5, e0))) | (~ (e6 = op(e6, e4)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e0, e4))) | (~ (e6 = op(e4, e6)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e4, e0))) | (~ (e5 = op(e6, e4)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e0, e4))) | (~ (e5 = op(e4, e6)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e4, e0))) | (~ (e4 = op(e6, e4)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | (~ (e4 = op(e4, e6)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | (~ (e3 = op(e6, e4)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | (~ (e3 = op(e4, e6)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | (~ (e2 = op(e6, e4)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | (~ (e2 = op(e4, e6)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | (~ (e1 = op(e6, e4)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | (~ (e1 = op(e4, e6)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | (~ (e0 = op(e6, e4)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | (~ (e0 = op(e4, e6)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e6 = op(e6, e3)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e0, e3))) | (~ (e6 = op(e3, e6)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e3, e0))) | (~ (e5 = op(e6, e3)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e0, e3))) | (~ (e5 = op(e3, e6)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e3, e0))) | (~ (e4 = op(e6, e3)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | (~ (e4 = op(e3, e6)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | (~ (e3 = op(e6, e3)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e6)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e6, e3)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e6)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e6, e3)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e6)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e6, e3)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e6)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e6 = op(e6, e2)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e0, e2))) | (~ (e6 = op(e2, e6)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e2, e0))) | (~ (e5 = op(e6, e2)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e0, e2))) | (~ (e5 = op(e2, e6)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e2, e0))) | (~ (e4 = op(e6, e2)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | (~ (e4 = op(e2, e6)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | (~ (e3 = op(e6, e2)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e6)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e6, e2)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e6)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e6, e2)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e6)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e6, e2)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e6)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e6 = op(e6, e1)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e0, e1))) | (~ (e6 = op(e1, e6)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e1, e0))) | (~ (e5 = op(e6, e1)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e0, e1))) | (~ (e5 = op(e1, e6)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e1, e0))) | (~ (e4 = op(e6, e1)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | (~ (e4 = op(e1, e6)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | (~ (e3 = op(e6, e1)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e6)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e6, e1)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e6)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e6, e1)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e6)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e6, e1)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e6)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e6 = op(e6, e0)) & ~ (e6 = op(e5, e0)) & ~ (e6 = op(e4, e0)) & ~ (e6 = op(e3, e0)) & ~ (e6 = op(e2, e0)) & ~ (e6 = op(e1, e0)) & ~ (e6 = op(e0, e0))) | (~ (e6 = op(e0, e6)) & ~ (e6 = op(e0, e5)) & ~ (e6 = op(e0, e4)) & ~ (e6 = op(e0, e3)) & ~ (e6 = op(e0, e2)) & ~ (e6 = op(e0, e1)) & ~ (e6 = op(e0, e0))) | (~ (e5 = op(e6, e0)) & ~ (e5 = op(e5, e0)) & ~ (e5 = op(e4, e0)) & ~ (e5 = op(e3, e0)) & ~ (e5 = op(e2, e0)) & ~ (e5 = op(e1, e0)) & ~ (e5 = op(e0, e0))) | (~ (e5 = op(e0, e6)) & ~ (e5 = op(e0, e5)) & ~ (e5 = op(e0, e4)) & ~ (e5 = op(e0, e3)) & ~ (e5 = op(e0, e2)) & ~ (e5 = op(e0, e1)) & ~ (e5 = op(e0, e0))) | (~ (e4 = op(e6, e0)) & ~ (e4 = op(e5, e0)) & ~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | (~ (e4 = op(e0, e6)) & ~ (e4 = op(e0, e5)) & ~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | (~ (e3 = op(e6, e0)) & ~ (e3 = op(e5, e0)) & ~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e6)) & ~ (e3 = op(e0, e5)) & ~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e6, e0)) & ~ (e2 = op(e5, e0)) & ~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e6)) & ~ (e2 = op(e0, e5)) & ~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e6, e0)) & ~ (e1 = op(e5, e0)) & ~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e6)) & ~ (e1 = op(e0, e5)) & ~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e6, e0)) & ~ (e0 = op(e5, e0)) & ~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e6)) & ~ (e0 = op(e0, e5)) & ~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e6 = op(e6, e6)) & ~ (e5 = op(e6, e6)) & ~ (e4 = op(e6, e6)) & ~ (e3 = op(e6, e6)) & ~ (e2 = op(e6, e6)) & ~ (e1 = op(e6, e6)) & ~ (e0 = op(e6, e6))) | (~ (e6 = op(e6, e5)) & ~ (e5 = op(e6, e5)) & ~ (e4 = op(e6, e5)) & ~ (e3 = op(e6, e5)) & ~ (e2 = op(e6, e5)) & ~ (e1 = op(e6, e5)) & ~ (e0 = op(e6, e5))) | (~ (e6 = op(e6, e4)) & ~ (e5 = op(e6, e4)) & ~ (e4 = op(e6, e4)) & ~ (e3 = op(e6, e4)) & ~ (e2 = op(e6, e4)) & ~ (e1 = op(e6, e4)) & ~ (e0 = op(e6, e4))) | (~ (e6 = op(e6, e3)) & ~ (e5 = op(e6, e3)) & ~ (e4 = op(e6, e3)) & ~ (e3 = op(e6, e3)) & ~ (e2 = op(e6, e3)) & ~ (e1 = op(e6, e3)) & ~ (e0 = op(e6, e3))) | (~ (e6 = op(e6, e2)) & ~ (e5 = op(e6, e2)) & ~ (e4 = op(e6, e2)) & ~ (e3 = op(e6, e2)) & ~ (e2 = op(e6, e2)) & ~ (e1 = op(e6, e2)) & ~ (e0 = op(e6, e2))) | (~ (e6 = op(e6, e1)) & ~ (e5 = op(e6, e1)) & ~ (e4 = op(e6, e1)) & ~ (e3 = op(e6, e1)) & ~ (e2 = op(e6, e1)) & ~ (e1 = op(e6, e1)) & ~ (e0 = op(e6, e1))) | (~ (e6 = op(e6, e0)) & ~ (e5 = op(e6, e0)) & ~ (e4 = op(e6, e0)) & ~ (e3 = op(e6, e0)) & ~ (e2 = op(e6, e0)) & ~ (e1 = op(e6, e0)) & ~ (e0 = op(e6, e0))) | (~ (e6 = op(e5, e6)) & ~ (e5 = op(e5, e6)) & ~ (e4 = op(e5, e6)) & ~ (e3 = op(e5, e6)) & ~ (e2 = op(e5, e6)) & ~ (e1 = op(e5, e6)) & ~ (e0 = op(e5, e6))) | (~ (e6 = op(e5, e5)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e5, e5)) & ~ (e3 = op(e5, e5)) & ~ (e2 = op(e5, e5)) & ~ (e1 = op(e5, e5)) & ~ (e0 = op(e5, e5))) | (~ (e6 = op(e5, e4)) & ~ (e5 = op(e5, e4)) & ~ (e4 = op(e5, e4)) & ~ (e3 = op(e5, e4)) & ~ (e2 = op(e5, e4)) & ~ (e1 = op(e5, e4)) & ~ (e0 = op(e5, e4))) | (~ (e6 = op(e5, e3)) & ~ (e5 = op(e5, e3)) & ~ (e4 = op(e5, e3)) & ~ (e3 = op(e5, e3)) & ~ (e2 = op(e5, e3)) & ~ (e1 = op(e5, e3)) & ~ (e0 = op(e5, e3))) | (~ (e6 = op(e5, e2)) & ~ (e5 = op(e5, e2)) & ~ (e4 = op(e5, e2)) & ~ (e3 = op(e5, e2)) & ~ (e2 = op(e5, e2)) & ~ (e1 = op(e5, e2)) & ~ (e0 = op(e5, e2))) | (~ (e6 = op(e5, e1)) & ~ (e5 = op(e5, e1)) & ~ (e4 = op(e5, e1)) & ~ (e3 = op(e5, e1)) & ~ (e2 = op(e5, e1)) & ~ (e1 = op(e5, e1)) & ~ (e0 = op(e5, e1))) | (~ (e6 = op(e5, e0)) & ~ (e5 = op(e5, e0)) & ~ (e4 = op(e5, e0)) & ~ (e3 = op(e5, e0)) & ~ (e2 = op(e5, e0)) & ~ (e1 = op(e5, e0)) & ~ (e0 = op(e5, e0))) | (~ (e6 = op(e4, e6)) & ~ (e5 = op(e4, e6)) & ~ (e4 = op(e4, e6)) & ~ (e3 = op(e4, e6)) & ~ (e2 = op(e4, e6)) & ~ (e1 = op(e4, e6)) & ~ (e0 = op(e4, e6))) | (~ (e6 = op(e4, e5)) & ~ (e5 = op(e4, e5)) & ~ (e4 = op(e4, e5)) & ~ (e3 = op(e4, e5)) & ~ (e2 = op(e4, e5)) & ~ (e1 = op(e4, e5)) & ~ (e0 = op(e4, e5))) | (~ (e6 = op(e4, e4)) & ~ (e5 = op(e4, e4)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | (~ (e6 = op(e4, e3)) & ~ (e5 = op(e4, e3)) & ~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | (~ (e6 = op(e4, e2)) & ~ (e5 = op(e4, e2)) & ~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | (~ (e6 = op(e4, e1)) & ~ (e5 = op(e4, e1)) & ~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | (~ (e6 = op(e4, e0)) & ~ (e5 = op(e4, e0)) & ~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | (~ (e6 = op(e3, e6)) & ~ (e5 = op(e3, e6)) & ~ (e4 = op(e3, e6)) & ~ (e3 = op(e3, e6)) & ~ (e2 = op(e3, e6)) & ~ (e1 = op(e3, e6)) & ~ (e0 = op(e3, e6))) | (~ (e6 = op(e3, e5)) & ~ (e5 = op(e3, e5)) & ~ (e4 = op(e3, e5)) & ~ (e3 = op(e3, e5)) & ~ (e2 = op(e3, e5)) & ~ (e1 = op(e3, e5)) & ~ (e0 = op(e3, e5))) | (~ (e6 = op(e3, e4)) & ~ (e5 = op(e3, e4)) & ~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | (~ (e6 = op(e3, e3)) & ~ (e5 = op(e3, e3)) & ~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e6 = op(e3, e2)) & ~ (e5 = op(e3, e2)) & ~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e6 = op(e3, e1)) & ~ (e5 = op(e3, e1)) & ~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e6 = op(e3, e0)) & ~ (e5 = op(e3, e0)) & ~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e6 = op(e2, e6)) & ~ (e5 = op(e2, e6)) & ~ (e4 = op(e2, e6)) & ~ (e3 = op(e2, e6)) & ~ (e2 = op(e2, e6)) & ~ (e1 = op(e2, e6)) & ~ (e0 = op(e2, e6))) | (~ (e6 = op(e2, e5)) & ~ (e5 = op(e2, e5)) & ~ (e4 = op(e2, e5)) & ~ (e3 = op(e2, e5)) & ~ (e2 = op(e2, e5)) & ~ (e1 = op(e2, e5)) & ~ (e0 = op(e2, e5))) | (~ (e6 = op(e2, e4)) & ~ (e5 = op(e2, e4)) & ~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | (~ (e6 = op(e2, e3)) & ~ (e5 = op(e2, e3)) & ~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e6 = op(e2, e2)) & ~ (e5 = op(e2, e2)) & ~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e6 = op(e2, e1)) & ~ (e5 = op(e2, e1)) & ~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e6 = op(e2, e0)) & ~ (e5 = op(e2, e0)) & ~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e6 = op(e1, e6)) & ~ (e5 = op(e1, e6)) & ~ (e4 = op(e1, e6)) & ~ (e3 = op(e1, e6)) & ~ (e2 = op(e1, e6)) & ~ (e1 = op(e1, e6)) & ~ (e0 = op(e1, e6))) | (~ (e6 = op(e1, e5)) & ~ (e5 = op(e1, e5)) & ~ (e4 = op(e1, e5)) & ~ (e3 = op(e1, e5)) & ~ (e2 = op(e1, e5)) & ~ (e1 = op(e1, e5)) & ~ (e0 = op(e1, e5))) | (~ (e6 = op(e1, e4)) & ~ (e5 = op(e1, e4)) & ~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | (~ (e6 = op(e1, e3)) & ~ (e5 = op(e1, e3)) & ~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e6 = op(e1, e2)) & ~ (e5 = op(e1, e2)) & ~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e6 = op(e1, e1)) & ~ (e5 = op(e1, e1)) & ~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e6 = op(e1, e0)) & ~ (e5 = op(e1, e0)) & ~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e6 = op(e0, e6)) & ~ (e5 = op(e0, e6)) & ~ (e4 = op(e0, e6)) & ~ (e3 = op(e0, e6)) & ~ (e2 = op(e0, e6)) & ~ (e1 = op(e0, e6)) & ~ (e0 = op(e0, e6))) | (~ (e6 = op(e0, e5)) & ~ (e5 = op(e0, e5)) & ~ (e4 = op(e0, e5)) & ~ (e3 = op(e0, e5)) & ~ (e2 = op(e0, e5)) & ~ (e1 = op(e0, e5)) & ~ (e0 = op(e0, e5))) | (~ (e6 = op(e0, e4)) & ~ (e5 = op(e0, e4)) & ~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | (~ (e6 = op(e0, e3)) & ~ (e5 = op(e0, e3)) & ~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e6 = op(e0, e2)) & ~ (e5 = op(e0, e2)) & ~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e6 = op(e0, e1)) & ~ (e5 = op(e0, e1)) & ~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e6 = op(e0, e0)) & ~ (e5 = op(e0, e0)) & ~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ (e6 = op(e6, e6)) | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ~ ((e6 = op(op(op(e6, e6), e6), e6)) & (e6 = op(op(op(e5, e6), e5), e5)) & (e6 = op(op(op(e4, e6), e4), e4)) & (e6 = op(op(op(e3, e6), e3), e3)) & (e6 = op(op(op(e2, e6), e2), e2)) & (e6 = op(op(op(e1, e6), e1), e1)) & (e6 = op(op(op(e0, e6), e0), e0)) & (e5 = op(op(op(e6, e5), e6), e6)) & (e5 = op(op(op(e5, e5), e5), e5)) & (e5 = op(op(op(e4, e5), e4), e4)) & (e5 = op(op(op(e3, e5), e3), e3)) & (e5 = op(op(op(e2, e5), e2), e2)) & (e5 = op(op(op(e1, e5), e1), e1)) & (e5 = op(op(op(e0, e5), e0), e0)) & (e4 = op(op(op(e6, e4), e6), e6)) & (e4 = op(op(op(e5, e4), e5), e5)) & (e4 = op(op(op(e4, e4), e4), e4)) & (e4 = op(op(op(e3, e4), e3), e3)) & (e4 = op(op(op(e2, e4), e2), e2)) & (e4 = op(op(op(e1, e4), e1), e1)) & (e4 = op(op(op(e0, e4), e0), e0)) & (e3 = op(op(op(e6, e3), e6), e6)) & (e3 = op(op(op(e5, e3), e5), e5)) & (e3 = op(op(op(e4, e3), e4), e4)) & (e3 = op(op(op(e3, e3), e3), e3)) & (e3 = op(op(op(e2, e3), e2), e2)) & (e3 = op(op(op(e1, e3), e1), e1)) & (e3 = op(op(op(e0, e3), e0), e0)) & (e2 = op(op(op(e6, e2), e6), e6)) & (e2 = op(op(op(e5, e2), e5), e5)) & (e2 = op(op(op(e4, e2), e4), e4)) & (e2 = op(op(op(e3, e2), e3), e3)) & (e2 = op(op(op(e2, e2), e2), e2)) & (e2 = op(op(op(e1, e2), e1), e1)) & (e2 = op(op(op(e0, e2), e0), e0)) & (e1 = op(op(op(e6, e1), e6), e6)) & (e1 = op(op(op(e5, e1), e5), e5)) & (e1 = op(op(op(e4, e1), e4), e4)) & (e1 = op(op(op(e3, e1), e3), e3)) & (e1 = op(op(op(e2, e1), e2), e2)) & (e1 = op(op(op(e1, e1), e1), e1)) & (e1 = op(op(op(e0, e1), e0), e0)) & (e0 = op(op(op(e6, e0), e6), e6)) & (e0 = op(op(op(e5, e0), e5), e5)) & (e0 = op(op(op(e4, e0), e4), e4)) & (e0 = op(op(op(e3, e0), e3), e3)) & (e0 = op(op(op(e2, e0), e2), e2)) & (e0 = op(op(op(e1, e0), e1), e1)) & (e0 = op(op(op(e0, e0), e0), e0)) & ((e6 = op(e6, e6)) | (e6 = op(e5, e6)) | (e6 = op(e4, e6)) | (e6 = op(e3, e6)) | (e6 = op(e2, e6)) | (e6 = op(e1, e6)) | (e6 = op(e0, e6))) & ((e6 = op(e6, e6)) | (e6 = op(e6, e5)) | (e6 = op(e6, e4)) | (e6 = op(e6, e3)) | (e6 = op(e6, e2)) | (e6 = op(e6, e1)) | (e6 = op(e6, e0))) & ((e5 = op(e6, e6)) | (e5 = op(e5, e6)) | (e5 = op(e4, e6)) | (e5 = op(e3, e6)) | (e5 = op(e2, e6)) | (e5 = op(e1, e6)) | (e5 = op(e0, e6))) & ((e5 = op(e6, e6)) | (e5 = op(e6, e5)) | (e5 = op(e6, e4)) | (e5 = op(e6, e3)) | (e5 = op(e6, e2)) | (e5 = op(e6, e1)) | (e5 = op(e6, e0))) & ((e4 = op(e6, e6)) | (e4 = op(e5, e6)) | (e4 = op(e4, e6)) | (e4 = op(e3, e6)) | (e4 = op(e2, e6)) | (e4 = op(e1, e6)) | (e4 = op(e0, e6))) & ((e4 = op(e6, e6)) | (e4 = op(e6, e5)) | (e4 = op(e6, e4)) | (e4 = op(e6, e3)) | (e4 = op(e6, e2)) | (e4 = op(e6, e1)) | (e4 = op(e6, e0))) & ((e3 = op(e6, e6)) | (e3 = op(e5, e6)) | (e3 = op(e4, e6)) | (e3 = op(e3, e6)) | (e3 = op(e2, e6)) | (e3 = op(e1, e6)) | (e3 = op(e0, e6))) & ((e3 = op(e6, e6)) | (e3 = op(e6, e5)) | (e3 = op(e6, e4)) | (e3 = op(e6, e3)) | (e3 = op(e6, e2)) | (e3 = op(e6, e1)) | (e3 = op(e6, e0))) & ((e2 = op(e6, e6)) | (e2 = op(e5, e6)) | (e2 = op(e4, e6)) | (e2 = op(e3, e6)) | (e2 = op(e2, e6)) | (e2 = op(e1, e6)) | (e2 = op(e0, e6))) & ((e2 = op(e6, e6)) | (e2 = op(e6, e5)) | (e2 = op(e6, e4)) | (e2 = op(e6, e3)) | (e2 = op(e6, e2)) | (e2 = op(e6, e1)) | (e2 = op(e6, e0))) & ((e1 = op(e6, e6)) | (e1 = op(e5, e6)) | (e1 = op(e4, e6)) | (e1 = op(e3, e6)) | (e1 = op(e2, e6)) | (e1 = op(e1, e6)) | (e1 = op(e0, e6))) & ((e1 = op(e6, e6)) | (e1 = op(e6, e5)) | (e1 = op(e6, e4)) | (e1 = op(e6, e3)) | (e1 = op(e6, e2)) | (e1 = op(e6, e1)) | (e1 = op(e6, e0))) & ((e0 = op(e6, e6)) | (e0 = op(e5, e6)) | (e0 = op(e4, e6)) | (e0 = op(e3, e6)) | (e0 = op(e2, e6)) | (e0 = op(e1, e6)) | (e0 = op(e0, e6))) & ((e0 = op(e6, e6)) | (e0 = op(e6, e5)) | (e0 = op(e6, e4)) | (e0 = op(e6, e3)) | (e0 = op(e6, e2)) | (e0 = op(e6, e1)) | (e0 = op(e6, e0))) & ((e6 = op(e6, e5)) | (e6 = op(e5, e5)) | (e6 = op(e4, e5)) | (e6 = op(e3, e5)) | (e6 = op(e2, e5)) | (e6 = op(e1, e5)) | (e6 = op(e0, e5))) & ((e6 = op(e5, e6)) | (e6 = op(e5, e5)) | (e6 = op(e5, e4)) | (e6 = op(e5, e3)) | (e6 = op(e5, e2)) | (e6 = op(e5, e1)) | (e6 = op(e5, e0))) & ((e5 = op(e6, e5)) | (e5 = op(e5, e5)) | (e5 = op(e4, e5)) | (e5 = op(e3, e5)) | (e5 = op(e2, e5)) | (e5 = op(e1, e5)) | (e5 = op(e0, e5))) & ((e5 = op(e5, e6)) | (e5 = op(e5, e5)) | (e5 = op(e5, e4)) | (e5 = op(e5, e3)) | (e5 = op(e5, e2)) | (e5 = op(e5, e1)) | (e5 = op(e5, e0))) & ((e4 = op(e6, e5)) | (e4 = op(e5, e5)) | (e4 = op(e4, e5)) | (e4 = op(e3, e5)) | (e4 = op(e2, e5)) | (e4 = op(e1, e5)) | (e4 = op(e0, e5))) & ((e4 = op(e5, e6)) | (e4 = op(e5, e5)) | (e4 = op(e5, e4)) | (e4 = op(e5, e3)) | (e4 = op(e5, e2)) | (e4 = op(e5, e1)) | (e4 = op(e5, e0))) & ((e3 = op(e6, e5)) | (e3 = op(e5, e5)) | (e3 = op(e4, e5)) | (e3 = op(e3, e5)) | (e3 = op(e2, e5)) | (e3 = op(e1, e5)) | (e3 = op(e0, e5))) & ((e3 = op(e5, e6)) | (e3 = op(e5, e5)) | (e3 = op(e5, e4)) | (e3 = op(e5, e3)) | (e3 = op(e5, e2)) | (e3 = op(e5, e1)) | (e3 = op(e5, e0))) & ((e2 = op(e6, e5)) | (e2 = op(e5, e5)) | (e2 = op(e4, e5)) | (e2 = op(e3, e5)) | (e2 = op(e2, e5)) | (e2 = op(e1, e5)) | (e2 = op(e0, e5))) & ((e2 = op(e5, e6)) | (e2 = op(e5, e5)) | (e2 = op(e5, e4)) | (e2 = op(e5, e3)) | (e2 = op(e5, e2)) | (e2 = op(e5, e1)) | (e2 = op(e5, e0))) & ((e1 = op(e6, e5)) | (e1 = op(e5, e5)) | (e1 = op(e4, e5)) | (e1 = op(e3, e5)) | (e1 = op(e2, e5)) | (e1 = op(e1, e5)) | (e1 = op(e0, e5))) & ((e1 = op(e5, e6)) | (e1 = op(e5, e5)) | (e1 = op(e5, e4)) | (e1 = op(e5, e3)) | (e1 = op(e5, e2)) | (e1 = op(e5, e1)) | (e1 = op(e5, e0))) & ((e0 = op(e6, e5)) | (e0 = op(e5, e5)) | (e0 = op(e4, e5)) | (e0 = op(e3, e5)) | (e0 = op(e2, e5)) | (e0 = op(e1, e5)) | (e0 = op(e0, e5))) & ((e0 = op(e5, e6)) | (e0 = op(e5, e5)) | (e0 = op(e5, e4)) | (e0 = op(e5, e3)) | (e0 = op(e5, e2)) | (e0 = op(e5, e1)) | (e0 = op(e5, e0))) & ((e6 = op(e6, e4)) | (e6 = op(e5, e4)) | (e6 = op(e4, e4)) | (e6 = op(e3, e4)) | (e6 = op(e2, e4)) | (e6 = op(e1, e4)) | (e6 = op(e0, e4))) & ((e6 = op(e4, e6)) | (e6 = op(e4, e5)) | (e6 = op(e4, e4)) | (e6 = op(e4, e3)) | (e6 = op(e4, e2)) | (e6 = op(e4, e1)) | (e6 = op(e4, e0))) & ((e5 = op(e6, e4)) | (e5 = op(e5, e4)) | (e5 = op(e4, e4)) | (e5 = op(e3, e4)) | (e5 = op(e2, e4)) | (e5 = op(e1, e4)) | (e5 = op(e0, e4))) & ((e5 = op(e4, e6)) | (e5 = op(e4, e5)) | (e5 = op(e4, e4)) | (e5 = op(e4, e3)) | (e5 = op(e4, e2)) | (e5 = op(e4, e1)) | (e5 = op(e4, e0))) & ((e4 = op(e6, e4)) | (e4 = op(e5, e4)) | (e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e6)) | (e4 = op(e4, e5)) | (e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e6, e4)) | (e3 = op(e5, e4)) | (e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e6)) | (e3 = op(e4, e5)) | (e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e6, e4)) | (e2 = op(e5, e4)) | (e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e6)) | (e2 = op(e4, e5)) | (e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e6, e4)) | (e1 = op(e5, e4)) | (e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e6)) | (e1 = op(e4, e5)) | (e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e6, e4)) | (e0 = op(e5, e4)) | (e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e6)) | (e0 = op(e4, e5)) | (e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e6 = op(e6, e3)) | (e6 = op(e5, e3)) | (e6 = op(e4, e3)) | (e6 = op(e3, e3)) | (e6 = op(e2, e3)) | (e6 = op(e1, e3)) | (e6 = op(e0, e3))) & ((e6 = op(e3, e6)) | (e6 = op(e3, e5)) | (e6 = op(e3, e4)) | (e6 = op(e3, e3)) | (e6 = op(e3, e2)) | (e6 = op(e3, e1)) | (e6 = op(e3, e0))) & ((e5 = op(e6, e3)) | (e5 = op(e5, e3)) | (e5 = op(e4, e3)) | (e5 = op(e3, e3)) | (e5 = op(e2, e3)) | (e5 = op(e1, e3)) | (e5 = op(e0, e3))) & ((e5 = op(e3, e6)) | (e5 = op(e3, e5)) | (e5 = op(e3, e4)) | (e5 = op(e3, e3)) | (e5 = op(e3, e2)) | (e5 = op(e3, e1)) | (e5 = op(e3, e0))) & ((e4 = op(e6, e3)) | (e4 = op(e5, e3)) | (e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e6)) | (e4 = op(e3, e5)) | (e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e6, e3)) | (e3 = op(e5, e3)) | (e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e6)) | (e3 = op(e3, e5)) | (e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e6, e3)) | (e2 = op(e5, e3)) | (e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e6)) | (e2 = op(e3, e5)) | (e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e6, e3)) | (e1 = op(e5, e3)) | (e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e6)) | (e1 = op(e3, e5)) | (e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e6, e3)) | (e0 = op(e5, e3)) | (e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e6)) | (e0 = op(e3, e5)) | (e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e6 = op(e6, e2)) | (e6 = op(e5, e2)) | (e6 = op(e4, e2)) | (e6 = op(e3, e2)) | (e6 = op(e2, e2)) | (e6 = op(e1, e2)) | (e6 = op(e0, e2))) & ((e6 = op(e2, e6)) | (e6 = op(e2, e5)) | (e6 = op(e2, e4)) | (e6 = op(e2, e3)) | (e6 = op(e2, e2)) | (e6 = op(e2, e1)) | (e6 = op(e2, e0))) & ((e5 = op(e6, e2)) | (e5 = op(e5, e2)) | (e5 = op(e4, e2)) | (e5 = op(e3, e2)) | (e5 = op(e2, e2)) | (e5 = op(e1, e2)) | (e5 = op(e0, e2))) & ((e5 = op(e2, e6)) | (e5 = op(e2, e5)) | (e5 = op(e2, e4)) | (e5 = op(e2, e3)) | (e5 = op(e2, e2)) | (e5 = op(e2, e1)) | (e5 = op(e2, e0))) & ((e4 = op(e6, e2)) | (e4 = op(e5, e2)) | (e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e6)) | (e4 = op(e2, e5)) | (e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e6, e2)) | (e3 = op(e5, e2)) | (e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e6)) | (e3 = op(e2, e5)) | (e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e6, e2)) | (e2 = op(e5, e2)) | (e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e6)) | (e2 = op(e2, e5)) | (e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e6, e2)) | (e1 = op(e5, e2)) | (e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e6)) | (e1 = op(e2, e5)) | (e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e6, e2)) | (e0 = op(e5, e2)) | (e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e6)) | (e0 = op(e2, e5)) | (e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e6 = op(e6, e1)) | (e6 = op(e5, e1)) | (e6 = op(e4, e1)) | (e6 = op(e3, e1)) | (e6 = op(e2, e1)) | (e6 = op(e1, e1)) | (e6 = op(e0, e1))) & ((e6 = op(e1, e6)) | (e6 = op(e1, e5)) | (e6 = op(e1, e4)) | (e6 = op(e1, e3)) | (e6 = op(e1, e2)) | (e6 = op(e1, e1)) | (e6 = op(e1, e0))) & ((e5 = op(e6, e1)) | (e5 = op(e5, e1)) | (e5 = op(e4, e1)) | (e5 = op(e3, e1)) | (e5 = op(e2, e1)) | (e5 = op(e1, e1)) | (e5 = op(e0, e1))) & ((e5 = op(e1, e6)) | (e5 = op(e1, e5)) | (e5 = op(e1, e4)) | (e5 = op(e1, e3)) | (e5 = op(e1, e2)) | (e5 = op(e1, e1)) | (e5 = op(e1, e0))) & ((e4 = op(e6, e1)) | (e4 = op(e5, e1)) | (e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e6)) | (e4 = op(e1, e5)) | (e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e6, e1)) | (e3 = op(e5, e1)) | (e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e6)) | (e3 = op(e1, e5)) | (e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e6, e1)) | (e2 = op(e5, e1)) | (e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e6)) | (e2 = op(e1, e5)) | (e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e6, e1)) | (e1 = op(e5, e1)) | (e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e6)) | (e1 = op(e1, e5)) | (e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e6, e1)) | (e0 = op(e5, e1)) | (e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e6)) | (e0 = op(e1, e5)) | (e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e6 = op(e6, e0)) | (e6 = op(e5, e0)) | (e6 = op(e4, e0)) | (e6 = op(e3, e0)) | (e6 = op(e2, e0)) | (e6 = op(e1, e0)) | (e6 = op(e0, e0))) & ((e6 = op(e0, e6)) | (e6 = op(e0, e5)) | (e6 = op(e0, e4)) | (e6 = op(e0, e3)) | (e6 = op(e0, e2)) | (e6 = op(e0, e1)) | (e6 = op(e0, e0))) & ((e5 = op(e6, e0)) | (e5 = op(e5, e0)) | (e5 = op(e4, e0)) | (e5 = op(e3, e0)) | (e5 = op(e2, e0)) | (e5 = op(e1, e0)) | (e5 = op(e0, e0))) & ((e5 = op(e0, e6)) | (e5 = op(e0, e5)) | (e5 = op(e0, e4)) | (e5 = op(e0, e3)) | (e5 = op(e0, e2)) | (e5 = op(e0, e1)) | (e5 = op(e0, e0))) & ((e4 = op(e6, e0)) | (e4 = op(e5, e0)) | (e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e6)) | (e4 = op(e0, e5)) | (e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e6, e0)) | (e3 = op(e5, e0)) | (e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e6)) | (e3 = op(e0, e5)) | (e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e6, e0)) | (e2 = op(e5, e0)) | (e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e6)) | (e2 = op(e0, e5)) | (e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e6, e0)) | (e1 = op(e5, e0)) | (e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e6)) | (e1 = op(e0, e5)) | (e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e6, e0)) | (e0 = op(e5, e0)) | (e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e6)) | (e0 = op(e0, e5)) | (e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e6 = op(e6, e6)) | (e5 = op(e6, e6)) | (e4 = op(e6, e6)) | (e3 = op(e6, e6)) | (e2 = op(e6, e6)) | (e1 = op(e6, e6)) | (e0 = op(e6, e6))) & ((e6 = op(e6, e5)) | (e5 = op(e6, e5)) | (e4 = op(e6, e5)) | (e3 = op(e6, e5)) | (e2 = op(e6, e5)) | (e1 = op(e6, e5)) | (e0 = op(e6, e5))) & ((e6 = op(e6, e4)) | (e5 = op(e6, e4)) | (e4 = op(e6, e4)) | (e3 = op(e6, e4)) | (e2 = op(e6, e4)) | (e1 = op(e6, e4)) | (e0 = op(e6, e4))) & ((e6 = op(e6, e3)) | (e5 = op(e6, e3)) | (e4 = op(e6, e3)) | (e3 = op(e6, e3)) | (e2 = op(e6, e3)) | (e1 = op(e6, e3)) | (e0 = op(e6, e3))) & ((e6 = op(e6, e2)) | (e5 = op(e6, e2)) | (e4 = op(e6, e2)) | (e3 = op(e6, e2)) | (e2 = op(e6, e2)) | (e1 = op(e6, e2)) | (e0 = op(e6, e2))) & ((e6 = op(e6, e1)) | (e5 = op(e6, e1)) | (e4 = op(e6, e1)) | (e3 = op(e6, e1)) | (e2 = op(e6, e1)) | (e1 = op(e6, e1)) | (e0 = op(e6, e1))) & ((e6 = op(e6, e0)) | (e5 = op(e6, e0)) | (e4 = op(e6, e0)) | (e3 = op(e6, e0)) | (e2 = op(e6, e0)) | (e1 = op(e6, e0)) | (e0 = op(e6, e0))) & ((e6 = op(e5, e6)) | (e5 = op(e5, e6)) | (e4 = op(e5, e6)) | (e3 = op(e5, e6)) | (e2 = op(e5, e6)) | (e1 = op(e5, e6)) | (e0 = op(e5, e6))) & ((e6 = op(e5, e5)) | (e5 = op(e5, e5)) | (e4 = op(e5, e5)) | (e3 = op(e5, e5)) | (e2 = op(e5, e5)) | (e1 = op(e5, e5)) | (e0 = op(e5, e5))) & ((e6 = op(e5, e4)) | (e5 = op(e5, e4)) | (e4 = op(e5, e4)) | (e3 = op(e5, e4)) | (e2 = op(e5, e4)) | (e1 = op(e5, e4)) | (e0 = op(e5, e4))) & ((e6 = op(e5, e3)) | (e5 = op(e5, e3)) | (e4 = op(e5, e3)) | (e3 = op(e5, e3)) | (e2 = op(e5, e3)) | (e1 = op(e5, e3)) | (e0 = op(e5, e3))) & ((e6 = op(e5, e2)) | (e5 = op(e5, e2)) | (e4 = op(e5, e2)) | (e3 = op(e5, e2)) | (e2 = op(e5, e2)) | (e1 = op(e5, e2)) | (e0 = op(e5, e2))) & ((e6 = op(e5, e1)) | (e5 = op(e5, e1)) | (e4 = op(e5, e1)) | (e3 = op(e5, e1)) | (e2 = op(e5, e1)) | (e1 = op(e5, e1)) | (e0 = op(e5, e1))) & ((e6 = op(e5, e0)) | (e5 = op(e5, e0)) | (e4 = op(e5, e0)) | (e3 = op(e5, e0)) | (e2 = op(e5, e0)) | (e1 = op(e5, e0)) | (e0 = op(e5, e0))) & ((e6 = op(e4, e6)) | (e5 = op(e4, e6)) | (e4 = op(e4, e6)) | (e3 = op(e4, e6)) | (e2 = op(e4, e6)) | (e1 = op(e4, e6)) | (e0 = op(e4, e6))) & ((e6 = op(e4, e5)) | (e5 = op(e4, e5)) | (e4 = op(e4, e5)) | (e3 = op(e4, e5)) | (e2 = op(e4, e5)) | (e1 = op(e4, e5)) | (e0 = op(e4, e5))) & ((e6 = op(e4, e4)) | (e5 = op(e4, e4)) | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e6 = op(e4, e3)) | (e5 = op(e4, e3)) | (e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e6 = op(e4, e2)) | (e5 = op(e4, e2)) | (e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e6 = op(e4, e1)) | (e5 = op(e4, e1)) | (e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e6 = op(e4, e0)) | (e5 = op(e4, e0)) | (e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e6 = op(e3, e6)) | (e5 = op(e3, e6)) | (e4 = op(e3, e6)) | (e3 = op(e3, e6)) | (e2 = op(e3, e6)) | (e1 = op(e3, e6)) | (e0 = op(e3, e6))) & ((e6 = op(e3, e5)) | (e5 = op(e3, e5)) | (e4 = op(e3, e5)) | (e3 = op(e3, e5)) | (e2 = op(e3, e5)) | (e1 = op(e3, e5)) | (e0 = op(e3, e5))) & ((e6 = op(e3, e4)) | (e5 = op(e3, e4)) | (e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e6 = op(e3, e3)) | (e5 = op(e3, e3)) | (e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e6 = op(e3, e2)) | (e5 = op(e3, e2)) | (e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e6 = op(e3, e1)) | (e5 = op(e3, e1)) | (e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e6 = op(e3, e0)) | (e5 = op(e3, e0)) | (e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e6 = op(e2, e6)) | (e5 = op(e2, e6)) | (e4 = op(e2, e6)) | (e3 = op(e2, e6)) | (e2 = op(e2, e6)) | (e1 = op(e2, e6)) | (e0 = op(e2, e6))) & ((e6 = op(e2, e5)) | (e5 = op(e2, e5)) | (e4 = op(e2, e5)) | (e3 = op(e2, e5)) | (e2 = op(e2, e5)) | (e1 = op(e2, e5)) | (e0 = op(e2, e5))) & ((e6 = op(e2, e4)) | (e5 = op(e2, e4)) | (e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e6 = op(e2, e3)) | (e5 = op(e2, e3)) | (e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e6 = op(e2, e2)) | (e5 = op(e2, e2)) | (e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e6 = op(e2, e1)) | (e5 = op(e2, e1)) | (e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e6 = op(e2, e0)) | (e5 = op(e2, e0)) | (e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e6 = op(e1, e6)) | (e5 = op(e1, e6)) | (e4 = op(e1, e6)) | (e3 = op(e1, e6)) | (e2 = op(e1, e6)) | (e1 = op(e1, e6)) | (e0 = op(e1, e6))) & ((e6 = op(e1, e5)) | (e5 = op(e1, e5)) | (e4 = op(e1, e5)) | (e3 = op(e1, e5)) | (e2 = op(e1, e5)) | (e1 = op(e1, e5)) | (e0 = op(e1, e5))) & ((e6 = op(e1, e4)) | (e5 = op(e1, e4)) | (e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e6 = op(e1, e3)) | (e5 = op(e1, e3)) | (e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e6 = op(e1, e2)) | (e5 = op(e1, e2)) | (e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e6 = op(e1, e1)) | (e5 = op(e1, e1)) | (e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e6 = op(e1, e0)) | (e5 = op(e1, e0)) | (e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e6 = op(e0, e6)) | (e5 = op(e0, e6)) | (e4 = op(e0, e6)) | (e3 = op(e0, e6)) | (e2 = op(e0, e6)) | (e1 = op(e0, e6)) | (e0 = op(e0, e6))) & ((e6 = op(e0, e5)) | (e5 = op(e0, e5)) | (e4 = op(e0, e5)) | (e3 = op(e0, e5)) | (e2 = op(e0, e5)) | (e1 = op(e0, e5)) | (e0 = op(e0, e5))) & ((e6 = op(e0, e4)) | (e5 = op(e0, e4)) | (e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e6 = op(e0, e3)) | (e5 = op(e0, e3)) | (e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e6 = op(e0, e2)) | (e5 = op(e0, e2)) | (e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e6 = op(e0, e1)) | (e5 = op(e0, e1)) | (e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e6 = op(e0, e0)) | (e5 = op(e0, e0)) | (e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & (e6 = op(e6, e6)) & (e5 = op(e5, e5)) & (e4 = op(e4, e4)) & (e3 = op(e3, e3)) & (e2 = op(e2, e2)) & (e1 = op(e1, e1)) & (e0 = op(e0, e0))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ((e6 = op(op(op(e6, e6), e6), e6)) & (e6 = op(op(op(e5, e6), e5), e5)) & (e6 = op(op(op(e4, e6), e4), e4)) & (e6 = op(op(op(e3, e6), e3), e3)) & (e6 = op(op(op(e2, e6), e2), e2)) & (e6 = op(op(op(e1, e6), e1), e1)) & (e6 = op(op(op(e0, e6), e0), e0)) & (e5 = op(op(op(e6, e5), e6), e6)) & (e5 = op(op(op(e5, e5), e5), e5)) & (e5 = op(op(op(e4, e5), e4), e4)) & (e5 = op(op(op(e3, e5), e3), e3)) & (e5 = op(op(op(e2, e5), e2), e2)) & (e5 = op(op(op(e1, e5), e1), e1)) & (e5 = op(op(op(e0, e5), e0), e0)) & (e4 = op(op(op(e6, e4), e6), e6)) & (e4 = op(op(op(e5, e4), e5), e5)) & (e4 = op(op(op(e4, e4), e4), e4)) & (e4 = op(op(op(e3, e4), e3), e3)) & (e4 = op(op(op(e2, e4), e2), e2)) & (e4 = op(op(op(e1, e4), e1), e1)) & (e4 = op(op(op(e0, e4), e0), e0)) & (e3 = op(op(op(e6, e3), e6), e6)) & (e3 = op(op(op(e5, e3), e5), e5)) & (e3 = op(op(op(e4, e3), e4), e4)) & (e3 = op(op(op(e3, e3), e3), e3)) & (e3 = op(op(op(e2, e3), e2), e2)) & (e3 = op(op(op(e1, e3), e1), e1)) & (e3 = op(op(op(e0, e3), e0), e0)) & (e2 = op(op(op(e6, e2), e6), e6)) & (e2 = op(op(op(e5, e2), e5), e5)) & (e2 = op(op(op(e4, e2), e4), e4)) & (e2 = op(op(op(e3, e2), e3), e3)) & (e2 = op(op(op(e2, e2), e2), e2)) & (e2 = op(op(op(e1, e2), e1), e1)) & (e2 = op(op(op(e0, e2), e0), e0)) & (e1 = op(op(op(e6, e1), e6), e6)) & (e1 = op(op(op(e5, e1), e5), e5)) & (e1 = op(op(op(e4, e1), e4), e4)) & (e1 = op(op(op(e3, e1), e3), e3)) & (e1 = op(op(op(e2, e1), e2), e2)) & (e1 = op(op(op(e1, e1), e1), e1)) & (e1 = op(op(op(e0, e1), e0), e0)) & (e0 = op(op(op(e6, e0), e6), e6)) & (e0 = op(op(op(e5, e0), e5), e5)) & (e0 = op(op(op(e4, e0), e4), e4)) & (e0 = op(op(op(e3, e0), e3), e3)) & (e0 = op(op(op(e2, e0), e2), e2)) & (e0 = op(op(op(e1, e0), e1), e1)) & (e0 = op(op(op(e0, e0), e0), e0)) & ((e6 = op(e6, e6)) | (e6 = op(e5, e6)) | (e6 = op(e4, e6)) | (e6 = op(e3, e6)) | (e6 = op(e2, e6)) | (e6 = op(e1, e6)) | (e6 = op(e0, e6))) & ((e6 = op(e6, e6)) | (e6 = op(e6, e5)) | (e6 = op(e6, e4)) | (e6 = op(e6, e3)) | (e6 = op(e6, e2)) | (e6 = op(e6, e1)) | (e6 = op(e6, e0))) & ((e5 = op(e6, e6)) | (e5 = op(e5, e6)) | (e5 = op(e4, e6)) | (e5 = op(e3, e6)) | (e5 = op(e2, e6)) | (e5 = op(e1, e6)) | (e5 = op(e0, e6))) & ((e5 = op(e6, e6)) | (e5 = op(e6, e5)) | (e5 = op(e6, e4)) | (e5 = op(e6, e3)) | (e5 = op(e6, e2)) | (e5 = op(e6, e1)) | (e5 = op(e6, e0))) & ((e4 = op(e6, e6)) | (e4 = op(e5, e6)) | (e4 = op(e4, e6)) | (e4 = op(e3, e6)) | (e4 = op(e2, e6)) | (e4 = op(e1, e6)) | (e4 = op(e0, e6))) & ((e4 = op(e6, e6)) | (e4 = op(e6, e5)) | (e4 = op(e6, e4)) | (e4 = op(e6, e3)) | (e4 = op(e6, e2)) | (e4 = op(e6, e1)) | (e4 = op(e6, e0))) & ((e3 = op(e6, e6)) | (e3 = op(e5, e6)) | (e3 = op(e4, e6)) | (e3 = op(e3, e6)) | (e3 = op(e2, e6)) | (e3 = op(e1, e6)) | (e3 = op(e0, e6))) & ((e3 = op(e6, e6)) | (e3 = op(e6, e5)) | (e3 = op(e6, e4)) | (e3 = op(e6, e3)) | (e3 = op(e6, e2)) | (e3 = op(e6, e1)) | (e3 = op(e6, e0))) & ((e2 = op(e6, e6)) | (e2 = op(e5, e6)) | (e2 = op(e4, e6)) | (e2 = op(e3, e6)) | (e2 = op(e2, e6)) | (e2 = op(e1, e6)) | (e2 = op(e0, e6))) & ((e2 = op(e6, e6)) | (e2 = op(e6, e5)) | (e2 = op(e6, e4)) | (e2 = op(e6, e3)) | (e2 = op(e6, e2)) | (e2 = op(e6, e1)) | (e2 = op(e6, e0))) & ((e1 = op(e6, e6)) | (e1 = op(e5, e6)) | (e1 = op(e4, e6)) | (e1 = op(e3, e6)) | (e1 = op(e2, e6)) | (e1 = op(e1, e6)) | (e1 = op(e0, e6))) & ((e1 = op(e6, e6)) | (e1 = op(e6, e5)) | (e1 = op(e6, e4)) | (e1 = op(e6, e3)) | (e1 = op(e6, e2)) | (e1 = op(e6, e1)) | (e1 = op(e6, e0))) & ((e0 = op(e6, e6)) | (e0 = op(e5, e6)) | (e0 = op(e4, e6)) | (e0 = op(e3, e6)) | (e0 = op(e2, e6)) | (e0 = op(e1, e6)) | (e0 = op(e0, e6))) & ((e0 = op(e6, e6)) | (e0 = op(e6, e5)) | (e0 = op(e6, e4)) | (e0 = op(e6, e3)) | (e0 = op(e6, e2)) | (e0 = op(e6, e1)) | (e0 = op(e6, e0))) & ((e6 = op(e6, e5)) | (e6 = op(e5, e5)) | (e6 = op(e4, e5)) | (e6 = op(e3, e5)) | (e6 = op(e2, e5)) | (e6 = op(e1, e5)) | (e6 = op(e0, e5))) & ((e6 = op(e5, e6)) | (e6 = op(e5, e5)) | (e6 = op(e5, e4)) | (e6 = op(e5, e3)) | (e6 = op(e5, e2)) | (e6 = op(e5, e1)) | (e6 = op(e5, e0))) & ((e5 = op(e6, e5)) | (e5 = op(e5, e5)) | (e5 = op(e4, e5)) | (e5 = op(e3, e5)) | (e5 = op(e2, e5)) | (e5 = op(e1, e5)) | (e5 = op(e0, e5))) & ((e5 = op(e5, e6)) | (e5 = op(e5, e5)) | (e5 = op(e5, e4)) | (e5 = op(e5, e3)) | (e5 = op(e5, e2)) | (e5 = op(e5, e1)) | (e5 = op(e5, e0))) & ((e4 = op(e6, e5)) | (e4 = op(e5, e5)) | (e4 = op(e4, e5)) | (e4 = op(e3, e5)) | (e4 = op(e2, e5)) | (e4 = op(e1, e5)) | (e4 = op(e0, e5))) & ((e4 = op(e5, e6)) | (e4 = op(e5, e5)) | (e4 = op(e5, e4)) | (e4 = op(e5, e3)) | (e4 = op(e5, e2)) | (e4 = op(e5, e1)) | (e4 = op(e5, e0))) & ((e3 = op(e6, e5)) | (e3 = op(e5, e5)) | (e3 = op(e4, e5)) | (e3 = op(e3, e5)) | (e3 = op(e2, e5)) | (e3 = op(e1, e5)) | (e3 = op(e0, e5))) & ((e3 = op(e5, e6)) | (e3 = op(e5, e5)) | (e3 = op(e5, e4)) | (e3 = op(e5, e3)) | (e3 = op(e5, e2)) | (e3 = op(e5, e1)) | (e3 = op(e5, e0))) & ((e2 = op(e6, e5)) | (e2 = op(e5, e5)) | (e2 = op(e4, e5)) | (e2 = op(e3, e5)) | (e2 = op(e2, e5)) | (e2 = op(e1, e5)) | (e2 = op(e0, e5))) & ((e2 = op(e5, e6)) | (e2 = op(e5, e5)) | (e2 = op(e5, e4)) | (e2 = op(e5, e3)) | (e2 = op(e5, e2)) | (e2 = op(e5, e1)) | (e2 = op(e5, e0))) & ((e1 = op(e6, e5)) | (e1 = op(e5, e5)) | (e1 = op(e4, e5)) | (e1 = op(e3, e5)) | (e1 = op(e2, e5)) | (e1 = op(e1, e5)) | (e1 = op(e0, e5))) & ((e1 = op(e5, e6)) | (e1 = op(e5, e5)) | (e1 = op(e5, e4)) | (e1 = op(e5, e3)) | (e1 = op(e5, e2)) | (e1 = op(e5, e1)) | (e1 = op(e5, e0))) & ((e0 = op(e6, e5)) | (e0 = op(e5, e5)) | (e0 = op(e4, e5)) | (e0 = op(e3, e5)) | (e0 = op(e2, e5)) | (e0 = op(e1, e5)) | (e0 = op(e0, e5))) & ((e0 = op(e5, e6)) | (e0 = op(e5, e5)) | (e0 = op(e5, e4)) | (e0 = op(e5, e3)) | (e0 = op(e5, e2)) | (e0 = op(e5, e1)) | (e0 = op(e5, e0))) & ((e6 = op(e6, e4)) | (e6 = op(e5, e4)) | (e6 = op(e4, e4)) | (e6 = op(e3, e4)) | (e6 = op(e2, e4)) | (e6 = op(e1, e4)) | (e6 = op(e0, e4))) & ((e6 = op(e4, e6)) | (e6 = op(e4, e5)) | (e6 = op(e4, e4)) | (e6 = op(e4, e3)) | (e6 = op(e4, e2)) | (e6 = op(e4, e1)) | (e6 = op(e4, e0))) & ((e5 = op(e6, e4)) | (e5 = op(e5, e4)) | (e5 = op(e4, e4)) | (e5 = op(e3, e4)) | (e5 = op(e2, e4)) | (e5 = op(e1, e4)) | (e5 = op(e0, e4))) & ((e5 = op(e4, e6)) | (e5 = op(e4, e5)) | (e5 = op(e4, e4)) | (e5 = op(e4, e3)) | (e5 = op(e4, e2)) | (e5 = op(e4, e1)) | (e5 = op(e4, e0))) & ((e4 = op(e6, e4)) | (e4 = op(e5, e4)) | (e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e6)) | (e4 = op(e4, e5)) | (e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e6, e4)) | (e3 = op(e5, e4)) | (e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e6)) | (e3 = op(e4, e5)) | (e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e6, e4)) | (e2 = op(e5, e4)) | (e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e6)) | (e2 = op(e4, e5)) | (e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e6, e4)) | (e1 = op(e5, e4)) | (e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e6)) | (e1 = op(e4, e5)) | (e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e6, e4)) | (e0 = op(e5, e4)) | (e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e6)) | (e0 = op(e4, e5)) | (e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e6 = op(e6, e3)) | (e6 = op(e5, e3)) | (e6 = op(e4, e3)) | (e6 = op(e3, e3)) | (e6 = op(e2, e3)) | (e6 = op(e1, e3)) | (e6 = op(e0, e3))) & ((e6 = op(e3, e6)) | (e6 = op(e3, e5)) | (e6 = op(e3, e4)) | (e6 = op(e3, e3)) | (e6 = op(e3, e2)) | (e6 = op(e3, e1)) | (e6 = op(e3, e0))) & ((e5 = op(e6, e3)) | (e5 = op(e5, e3)) | (e5 = op(e4, e3)) | (e5 = op(e3, e3)) | (e5 = op(e2, e3)) | (e5 = op(e1, e3)) | (e5 = op(e0, e3))) & ((e5 = op(e3, e6)) | (e5 = op(e3, e5)) | (e5 = op(e3, e4)) | (e5 = op(e3, e3)) | (e5 = op(e3, e2)) | (e5 = op(e3, e1)) | (e5 = op(e3, e0))) & ((e4 = op(e6, e3)) | (e4 = op(e5, e3)) | (e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e6)) | (e4 = op(e3, e5)) | (e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e6, e3)) | (e3 = op(e5, e3)) | (e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e6)) | (e3 = op(e3, e5)) | (e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e6, e3)) | (e2 = op(e5, e3)) | (e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e6)) | (e2 = op(e3, e5)) | (e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e6, e3)) | (e1 = op(e5, e3)) | (e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e6)) | (e1 = op(e3, e5)) | (e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e6, e3)) | (e0 = op(e5, e3)) | (e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e6)) | (e0 = op(e3, e5)) | (e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e6 = op(e6, e2)) | (e6 = op(e5, e2)) | (e6 = op(e4, e2)) | (e6 = op(e3, e2)) | (e6 = op(e2, e2)) | (e6 = op(e1, e2)) | (e6 = op(e0, e2))) & ((e6 = op(e2, e6)) | (e6 = op(e2, e5)) | (e6 = op(e2, e4)) | (e6 = op(e2, e3)) | (e6 = op(e2, e2)) | (e6 = op(e2, e1)) | (e6 = op(e2, e0))) & ((e5 = op(e6, e2)) | (e5 = op(e5, e2)) | (e5 = op(e4, e2)) | (e5 = op(e3, e2)) | (e5 = op(e2, e2)) | (e5 = op(e1, e2)) | (e5 = op(e0, e2))) & ((e5 = op(e2, e6)) | (e5 = op(e2, e5)) | (e5 = op(e2, e4)) | (e5 = op(e2, e3)) | (e5 = op(e2, e2)) | (e5 = op(e2, e1)) | (e5 = op(e2, e0))) & ((e4 = op(e6, e2)) | (e4 = op(e5, e2)) | (e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e6)) | (e4 = op(e2, e5)) | (e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e6, e2)) | (e3 = op(e5, e2)) | (e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e6)) | (e3 = op(e2, e5)) | (e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e6, e2)) | (e2 = op(e5, e2)) | (e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e6)) | (e2 = op(e2, e5)) | (e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e6, e2)) | (e1 = op(e5, e2)) | (e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e6)) | (e1 = op(e2, e5)) | (e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e6, e2)) | (e0 = op(e5, e2)) | (e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e6)) | (e0 = op(e2, e5)) | (e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e6 = op(e6, e1)) | (e6 = op(e5, e1)) | (e6 = op(e4, e1)) | (e6 = op(e3, e1)) | (e6 = op(e2, e1)) | (e6 = op(e1, e1)) | (e6 = op(e0, e1))) & ((e6 = op(e1, e6)) | (e6 = op(e1, e5)) | (e6 = op(e1, e4)) | (e6 = op(e1, e3)) | (e6 = op(e1, e2)) | (e6 = op(e1, e1)) | (e6 = op(e1, e0))) & ((e5 = op(e6, e1)) | (e5 = op(e5, e1)) | (e5 = op(e4, e1)) | (e5 = op(e3, e1)) | (e5 = op(e2, e1)) | (e5 = op(e1, e1)) | (e5 = op(e0, e1))) & ((e5 = op(e1, e6)) | (e5 = op(e1, e5)) | (e5 = op(e1, e4)) | (e5 = op(e1, e3)) | (e5 = op(e1, e2)) | (e5 = op(e1, e1)) | (e5 = op(e1, e0))) & ((e4 = op(e6, e1)) | (e4 = op(e5, e1)) | (e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e6)) | (e4 = op(e1, e5)) | (e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e6, e1)) | (e3 = op(e5, e1)) | (e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e6)) | (e3 = op(e1, e5)) | (e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e6, e1)) | (e2 = op(e5, e1)) | (e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e6)) | (e2 = op(e1, e5)) | (e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e6, e1)) | (e1 = op(e5, e1)) | (e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e6)) | (e1 = op(e1, e5)) | (e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e6, e1)) | (e0 = op(e5, e1)) | (e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e6)) | (e0 = op(e1, e5)) | (e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e6 = op(e6, e0)) | (e6 = op(e5, e0)) | (e6 = op(e4, e0)) | (e6 = op(e3, e0)) | (e6 = op(e2, e0)) | (e6 = op(e1, e0)) | (e6 = op(e0, e0))) & ((e6 = op(e0, e6)) | (e6 = op(e0, e5)) | (e6 = op(e0, e4)) | (e6 = op(e0, e3)) | (e6 = op(e0, e2)) | (e6 = op(e0, e1)) | (e6 = op(e0, e0))) & ((e5 = op(e6, e0)) | (e5 = op(e5, e0)) | (e5 = op(e4, e0)) | (e5 = op(e3, e0)) | (e5 = op(e2, e0)) | (e5 = op(e1, e0)) | (e5 = op(e0, e0))) & ((e5 = op(e0, e6)) | (e5 = op(e0, e5)) | (e5 = op(e0, e4)) | (e5 = op(e0, e3)) | (e5 = op(e0, e2)) | (e5 = op(e0, e1)) | (e5 = op(e0, e0))) & ((e4 = op(e6, e0)) | (e4 = op(e5, e0)) | (e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e6)) | (e4 = op(e0, e5)) | (e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e6, e0)) | (e3 = op(e5, e0)) | (e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e6)) | (e3 = op(e0, e5)) | (e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e6, e0)) | (e2 = op(e5, e0)) | (e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e6)) | (e2 = op(e0, e5)) | (e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e6, e0)) | (e1 = op(e5, e0)) | (e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e6)) | (e1 = op(e0, e5)) | (e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e6, e0)) | (e0 = op(e5, e0)) | (e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e6)) | (e0 = op(e0, e5)) | (e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e6 = op(e6, e6)) | (e5 = op(e6, e6)) | (e4 = op(e6, e6)) | (e3 = op(e6, e6)) | (e2 = op(e6, e6)) | (e1 = op(e6, e6)) | (e0 = op(e6, e6))) & ((e6 = op(e6, e5)) | (e5 = op(e6, e5)) | (e4 = op(e6, e5)) | (e3 = op(e6, e5)) | (e2 = op(e6, e5)) | (e1 = op(e6, e5)) | (e0 = op(e6, e5))) & ((e6 = op(e6, e4)) | (e5 = op(e6, e4)) | (e4 = op(e6, e4)) | (e3 = op(e6, e4)) | (e2 = op(e6, e4)) | (e1 = op(e6, e4)) | (e0 = op(e6, e4))) & ((e6 = op(e6, e3)) | (e5 = op(e6, e3)) | (e4 = op(e6, e3)) | (e3 = op(e6, e3)) | (e2 = op(e6, e3)) | (e1 = op(e6, e3)) | (e0 = op(e6, e3))) & ((e6 = op(e6, e2)) | (e5 = op(e6, e2)) | (e4 = op(e6, e2)) | (e3 = op(e6, e2)) | (e2 = op(e6, e2)) | (e1 = op(e6, e2)) | (e0 = op(e6, e2))) & ((e6 = op(e6, e1)) | (e5 = op(e6, e1)) | (e4 = op(e6, e1)) | (e3 = op(e6, e1)) | (e2 = op(e6, e1)) | (e1 = op(e6, e1)) | (e0 = op(e6, e1))) & ((e6 = op(e6, e0)) | (e5 = op(e6, e0)) | (e4 = op(e6, e0)) | (e3 = op(e6, e0)) | (e2 = op(e6, e0)) | (e1 = op(e6, e0)) | (e0 = op(e6, e0))) & ((e6 = op(e5, e6)) | (e5 = op(e5, e6)) | (e4 = op(e5, e6)) | (e3 = op(e5, e6)) | (e2 = op(e5, e6)) | (e1 = op(e5, e6)) | (e0 = op(e5, e6))) & ((e6 = op(e5, e5)) | (e5 = op(e5, e5)) | (e4 = op(e5, e5)) | (e3 = op(e5, e5)) | (e2 = op(e5, e5)) | (e1 = op(e5, e5)) | (e0 = op(e5, e5))) & ((e6 = op(e5, e4)) | (e5 = op(e5, e4)) | (e4 = op(e5, e4)) | (e3 = op(e5, e4)) | (e2 = op(e5, e4)) | (e1 = op(e5, e4)) | (e0 = op(e5, e4))) & ((e6 = op(e5, e3)) | (e5 = op(e5, e3)) | (e4 = op(e5, e3)) | (e3 = op(e5, e3)) | (e2 = op(e5, e3)) | (e1 = op(e5, e3)) | (e0 = op(e5, e3))) & ((e6 = op(e5, e2)) | (e5 = op(e5, e2)) | (e4 = op(e5, e2)) | (e3 = op(e5, e2)) | (e2 = op(e5, e2)) | (e1 = op(e5, e2)) | (e0 = op(e5, e2))) & ((e6 = op(e5, e1)) | (e5 = op(e5, e1)) | (e4 = op(e5, e1)) | (e3 = op(e5, e1)) | (e2 = op(e5, e1)) | (e1 = op(e5, e1)) | (e0 = op(e5, e1))) & ((e6 = op(e5, e0)) | (e5 = op(e5, e0)) | (e4 = op(e5, e0)) | (e3 = op(e5, e0)) | (e2 = op(e5, e0)) | (e1 = op(e5, e0)) | (e0 = op(e5, e0))) & ((e6 = op(e4, e6)) | (e5 = op(e4, e6)) | (e4 = op(e4, e6)) | (e3 = op(e4, e6)) | (e2 = op(e4, e6)) | (e1 = op(e4, e6)) | (e0 = op(e4, e6))) & ((e6 = op(e4, e5)) | (e5 = op(e4, e5)) | (e4 = op(e4, e5)) | (e3 = op(e4, e5)) | (e2 = op(e4, e5)) | (e1 = op(e4, e5)) | (e0 = op(e4, e5))) & ((e6 = op(e4, e4)) | (e5 = op(e4, e4)) | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e6 = op(e4, e3)) | (e5 = op(e4, e3)) | (e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e6 = op(e4, e2)) | (e5 = op(e4, e2)) | (e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e6 = op(e4, e1)) | (e5 = op(e4, e1)) | (e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e6 = op(e4, e0)) | (e5 = op(e4, e0)) | (e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e6 = op(e3, e6)) | (e5 = op(e3, e6)) | (e4 = op(e3, e6)) | (e3 = op(e3, e6)) | (e2 = op(e3, e6)) | (e1 = op(e3, e6)) | (e0 = op(e3, e6))) & ((e6 = op(e3, e5)) | (e5 = op(e3, e5)) | (e4 = op(e3, e5)) | (e3 = op(e3, e5)) | (e2 = op(e3, e5)) | (e1 = op(e3, e5)) | (e0 = op(e3, e5))) & ((e6 = op(e3, e4)) | (e5 = op(e3, e4)) | (e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e6 = op(e3, e3)) | (e5 = op(e3, e3)) | (e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e6 = op(e3, e2)) | (e5 = op(e3, e2)) | (e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e6 = op(e3, e1)) | (e5 = op(e3, e1)) | (e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e6 = op(e3, e0)) | (e5 = op(e3, e0)) | (e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e6 = op(e2, e6)) | (e5 = op(e2, e6)) | (e4 = op(e2, e6)) | (e3 = op(e2, e6)) | (e2 = op(e2, e6)) | (e1 = op(e2, e6)) | (e0 = op(e2, e6))) & ((e6 = op(e2, e5)) | (e5 = op(e2, e5)) | (e4 = op(e2, e5)) | (e3 = op(e2, e5)) | (e2 = op(e2, e5)) | (e1 = op(e2, e5)) | (e0 = op(e2, e5))) & ((e6 = op(e2, e4)) | (e5 = op(e2, e4)) | (e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e6 = op(e2, e3)) | (e5 = op(e2, e3)) | (e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e6 = op(e2, e2)) | (e5 = op(e2, e2)) | (e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e6 = op(e2, e1)) | (e5 = op(e2, e1)) | (e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e6 = op(e2, e0)) | (e5 = op(e2, e0)) | (e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e6 = op(e1, e6)) | (e5 = op(e1, e6)) | (e4 = op(e1, e6)) | (e3 = op(e1, e6)) | (e2 = op(e1, e6)) | (e1 = op(e1, e6)) | (e0 = op(e1, e6))) & ((e6 = op(e1, e5)) | (e5 = op(e1, e5)) | (e4 = op(e1, e5)) | (e3 = op(e1, e5)) | (e2 = op(e1, e5)) | (e1 = op(e1, e5)) | (e0 = op(e1, e5))) & ((e6 = op(e1, e4)) | (e5 = op(e1, e4)) | (e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e6 = op(e1, e3)) | (e5 = op(e1, e3)) | (e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e6 = op(e1, e2)) | (e5 = op(e1, e2)) | (e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e6 = op(e1, e1)) | (e5 = op(e1, e1)) | (e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e6 = op(e1, e0)) | (e5 = op(e1, e0)) | (e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e6 = op(e0, e6)) | (e5 = op(e0, e6)) | (e4 = op(e0, e6)) | (e3 = op(e0, e6)) | (e2 = op(e0, e6)) | (e1 = op(e0, e6)) | (e0 = op(e0, e6))) & ((e6 = op(e0, e5)) | (e5 = op(e0, e5)) | (e4 = op(e0, e5)) | (e3 = op(e0, e5)) | (e2 = op(e0, e5)) | (e1 = op(e0, e5)) | (e0 = op(e0, e5))) & ((e6 = op(e0, e4)) | (e5 = op(e0, e4)) | (e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e6 = op(e0, e3)) | (e5 = op(e0, e3)) | (e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e6 = op(e0, e2)) | (e5 = op(e0, e2)) | (e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e6 = op(e0, e1)) | (e5 = op(e0, e1)) | (e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e6 = op(e0, e0)) | (e5 = op(e0, e0)) | (e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & (e6 = op(e6, e6)) & (e5 = op(e5, e5)) & (e4 = op(e4, e4)) & (e3 = op(e3, e3)) & (e2 = op(e2, e2)) & (e1 = op(e1, e1)) & (e0 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG208+1.p', co1)).