fof(f4073, plain, $false, inference(subsumption_resolution, [], [f4072, f331])).
fof(f331, plain, (e0 = op(e1, e0)), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ((e5 = op(e6, e6)) & (e0 = op(e6, e5)) & (e1 = op(e6, e4)) & (e4 = op(e6, e3)) & (e2 = op(e6, e2)) & (e3 = op(e6, e1)) & (e6 = op(e6, e0)) & (e0 = op(e5, e6)) & (e1 = op(e5, e5)) & (e6 = op(e5, e4)) & (e3 = op(e5, e3)) & (e5 = op(e5, e2)) & (e2 = op(e5, e1)) & (e4 = op(e5, e0)) & (e1 = op(e4, e6)) & (e6 = op(e4, e5)) & (e4 = op(e4, e4)) & (e2 = op(e4, e3)) & (e0 = op(e4, e2)) & (e5 = op(e4, e1)) & (e3 = op(e4, e0)) & (e4 = op(e3, e6)) & (e3 = op(e3, e5)) & (e2 = op(e3, e4)) & (e0 = op(e3, e3)) & (e6 = op(e3, e2)) & (e1 = op(e3, e1)) & (e5 = op(e3, e0)) & (e2 = op(e2, e6)) & (e5 = op(e2, e5)) & (e0 = op(e2, e4)) & (e6 = op(e2, e3)) & (e3 = op(e2, e2)) & (e4 = op(e2, e1)) & (e1 = op(e2, e0)) & (e3 = op(e1, e6)) & (e2 = op(e1, e5)) & (e5 = op(e1, e4)) & (e1 = op(e1, e3)) & (e4 = op(e1, e2)) & (e6 = op(e1, e1)) & (e0 = op(e1, e0)) & (e6 = op(e0, e6)) & (e4 = op(e0, e5)) & (e3 = op(e0, e4)) & (e5 = op(e0, e3)) & (e1 = op(e0, e2)) & (e0 = op(e0, e1)) & (e2 = op(e0, e0))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG209+1.p', ax2)).
fof(f4072, plain, ~ (e0 = op(e1, e0)), inference(forward_demodulation, [], [f4071, f338])).
fof(f338, plain, (e1 = op(e2, e0)), inference(cnf_transformation, [], [f2])).
fof(f4071, plain, ~ (e0 = op(op(e2, e0), e0)), inference(forward_demodulation, [], [f4070, f324])).
fof(f324, plain, (e2 = op(e0, e0)), inference(cnf_transformation, [], [f2])).
fof(f4070, plain, ~ (e0 = op(op(op(e0, e0), e0), e0)), inference(subsumption_resolution, [], [f4069, f325])).
fof(f325, plain, (e0 = op(e0, e1)), inference(cnf_transformation, [], [f2])).
fof(f4069, plain, (~ (e0 = op(e0, e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4068, f325])).
fof(f4068, plain, (~ (e0 = op(op(e0, e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4067, f331])).
fof(f4067, plain, (~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4066, f354])).
fof(f354, plain, (e0 = op(e4, e2)), inference(cnf_transformation, [], [f2])).
fof(f4066, plain, (~ (e0 = op(e4, e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4065, f333])).
fof(f333, plain, (e4 = op(e1, e2)), inference(cnf_transformation, [], [f2])).
fof(f4065, plain, (~ (e0 = op(op(e1, e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4064, f338])).
fof(f4064, plain, (~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4063, f348])).
fof(f348, plain, (e0 = op(e3, e3)), inference(cnf_transformation, [], [f2])).
fof(f4063, plain, (~ (e0 = op(e3, e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4062, f362])).
fof(f362, plain, (e3 = op(e5, e3)), inference(cnf_transformation, [], [f2])).
fof(f4062, plain, (~ (e0 = op(op(e5, e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4061, f345])).
fof(f345, plain, (e5 = op(e3, e0)), inference(cnf_transformation, [], [f2])).
fof(f4061, plain, (~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4060, f342])).
fof(f342, plain, (e0 = op(e2, e4)), inference(cnf_transformation, [], [f2])).
fof(f4060, plain, (~ (e0 = op(e2, e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4059, f349])).
fof(f349, plain, (e2 = op(e3, e4)), inference(cnf_transformation, [], [f2])).
fof(f4059, plain, (~ (e0 = op(op(e3, e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4058, f352])).
fof(f352, plain, (e3 = op(e4, e0)), inference(cnf_transformation, [], [f2])).
fof(f4058, plain, (~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4057, f371])).
fof(f371, plain, (e0 = op(e6, e5)), inference(cnf_transformation, [], [f2])).
fof(f4057, plain, (~ (e0 = op(e6, e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4056, f357])).
fof(f357, plain, (e6 = op(e4, e5)), inference(cnf_transformation, [], [f2])).
fof(f4056, plain, (~ (e0 = op(op(e4, e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4055, f359])).
fof(f359, plain, (e4 = op(e5, e0)), inference(cnf_transformation, [], [f2])).
fof(f4055, plain, (~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4054, f365])).
fof(f365, plain, (e0 = op(e5, e6)), inference(cnf_transformation, [], [f2])).
fof(f4054, plain, (~ (e0 = op(e5, e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4053, f372])).
fof(f372, plain, (e5 = op(e6, e6)), inference(cnf_transformation, [], [f2])).
fof(f4053, plain, (~ (e0 = op(op(e6, e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4052, f366])).
fof(f366, plain, (e6 = op(e6, e0)), inference(cnf_transformation, [], [f2])).
fof(f4052, plain, (~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4051, f338])).
fof(f4051, plain, (~ (e1 = op(e2, e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4050, f324])).
fof(f4050, plain, (~ (e1 = op(op(e0, e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4049, f325])).
fof(f4049, plain, (~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4048, f346])).
fof(f346, plain, (e1 = op(e3, e1)), inference(cnf_transformation, [], [f2])).
fof(f4048, plain, (~ (e1 = op(e3, e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4047, f367])).
fof(f367, plain, (e3 = op(e6, e1)), inference(cnf_transformation, [], [f2])).
fof(f4047, plain, (~ (e1 = op(op(e6, e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4046, f332])).
fof(f332, plain, (e6 = op(e1, e1)), inference(cnf_transformation, [], [f2])).
fof(f4046, plain, (~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4045, f326])).
fof(f326, plain, (e1 = op(e0, e2)), inference(cnf_transformation, [], [f2])).
fof(f4045, plain, (~ (e1 = op(e0, e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4044, f354])).
fof(f4044, plain, (~ (e1 = op(op(e4, e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4043, f339])).
fof(f339, plain, (e4 = op(e2, e1)), inference(cnf_transformation, [], [f2])).
fof(f4043, plain, (~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4042, f334])).
fof(f334, plain, (e1 = op(e1, e3)), inference(cnf_transformation, [], [f2])).
fof(f4042, plain, (~ (e1 = op(e1, e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4041, f334])).
fof(f4041, plain, (~ (e1 = op(op(e1, e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4040, f346])).
fof(f4040, plain, (~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4039, f370])).
fof(f370, plain, (e1 = op(e6, e4)), inference(cnf_transformation, [], [f2])).
fof(f4039, plain, (~ (e1 = op(e6, e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4038, f363])).
fof(f363, plain, (e6 = op(e5, e4)), inference(cnf_transformation, [], [f2])).
fof(f4038, plain, (~ (e1 = op(op(e5, e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4037, f353])).
fof(f353, plain, (e5 = op(e4, e1)), inference(cnf_transformation, [], [f2])).
fof(f4037, plain, (~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4036, f364])).
fof(f364, plain, (e1 = op(e5, e5)), inference(cnf_transformation, [], [f2])).
fof(f4036, plain, (~ (e1 = op(e5, e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4035, f343])).
fof(f343, plain, (e5 = op(e2, e5)), inference(cnf_transformation, [], [f2])).
fof(f4035, plain, (~ (e1 = op(op(e2, e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4034, f360])).
fof(f360, plain, (e2 = op(e5, e1)), inference(cnf_transformation, [], [f2])).
fof(f4034, plain, (~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4033, f358])).
fof(f358, plain, (e1 = op(e4, e6)), inference(cnf_transformation, [], [f2])).
fof(f4033, plain, (~ (e1 = op(e4, e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4032, f351])).
fof(f351, plain, (e4 = op(e3, e6)), inference(cnf_transformation, [], [f2])).
fof(f4032, plain, (~ (e1 = op(op(e3, e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4031, f367])).
fof(f4031, plain, (~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4030, f324])).
fof(f4030, plain, (~ (e2 = op(e0, e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4029, f331])).
fof(f4029, plain, (~ (e2 = op(op(e1, e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4028, f326])).
fof(f4028, plain, (~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4027, f360])).
fof(f4027, plain, (~ (e2 = op(e5, e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4026, f353])).
fof(f4026, plain, (~ (e2 = op(op(e4, e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4025, f333])).
fof(f4025, plain, (~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4024, f368])).
fof(f368, plain, (e2 = op(e6, e2)), inference(cnf_transformation, [], [f2])).
fof(f4024, plain, (~ (e2 = op(e6, e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4023, f347])).
fof(f347, plain, (e6 = op(e3, e2)), inference(cnf_transformation, [], [f2])).
fof(f4023, plain, (~ (e2 = op(op(e3, e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4022, f340])).
fof(f340, plain, (e3 = op(e2, e2)), inference(cnf_transformation, [], [f2])).
fof(f4022, plain, (~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4021, f355])).
fof(f355, plain, (e2 = op(e4, e3)), inference(cnf_transformation, [], [f2])).
fof(f4021, plain, (~ (e2 = op(e4, e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4020, f369])).
fof(f369, plain, (e4 = op(e6, e3)), inference(cnf_transformation, [], [f2])).
fof(f4020, plain, (~ (e2 = op(op(e6, e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4019, f347])).
fof(f4019, plain, (~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4018, f349])).
fof(f4018, plain, (~ (e2 = op(e3, e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4017, f328])).
fof(f328, plain, (e3 = op(e0, e4)), inference(cnf_transformation, [], [f2])).
fof(f4017, plain, (~ (e2 = op(op(e0, e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4016, f354])).
fof(f4016, plain, (~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4015, f336])).
fof(f336, plain, (e2 = op(e1, e5)), inference(cnf_transformation, [], [f2])).
fof(f4015, plain, (~ (e2 = op(e1, e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4014, f364])).
fof(f4014, plain, (~ (e2 = op(op(e5, e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4013, f361])).
fof(f361, plain, (e5 = op(e5, e2)), inference(cnf_transformation, [], [f2])).
fof(f4013, plain, (~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4012, f344])).
fof(f344, plain, (e2 = op(e2, e6)), inference(cnf_transformation, [], [f2])).
fof(f4012, plain, (~ (e2 = op(e2, e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4011, f344])).
fof(f4011, plain, (~ (e2 = op(op(e2, e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4010, f368])).
fof(f4010, plain, (~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4009, f352])).
fof(f4009, plain, (~ (e3 = op(e4, e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4008, f359])).
fof(f4008, plain, (~ (e3 = op(op(e5, e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4007, f327])).
fof(f327, plain, (e5 = op(e0, e3)), inference(cnf_transformation, [], [f2])).
fof(f4007, plain, (~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4006, f367])).
fof(f4006, plain, (~ (e3 = op(e6, e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4005, f332])).
fof(f4005, plain, (~ (e3 = op(op(e1, e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4004, f334])).
fof(f4004, plain, (~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4003, f340])).
fof(f4003, plain, (~ (e3 = op(e2, e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4002, f368])).
fof(f4002, plain, (~ (e3 = op(op(e6, e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f4001, f341])).
fof(f341, plain, (e6 = op(e2, e3)), inference(cnf_transformation, [], [f2])).
fof(f4001, plain, (~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f4000, f362])).
fof(f4000, plain, (~ (e3 = op(e5, e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3999, f327])).
fof(f3999, plain, (~ (e3 = op(op(e0, e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3998, f348])).
fof(f3998, plain, (~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3997, f328])).
fof(f3997, plain, (~ (e3 = op(e0, e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3996, f342])).
fof(f3996, plain, (~ (e3 = op(op(e2, e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3995, f355])).
fof(f3995, plain, (~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3994, f350])).
fof(f350, plain, (e3 = op(e3, e5)), inference(cnf_transformation, [], [f2])).
fof(f3994, plain, (~ (e3 = op(e3, e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3993, f350])).
fof(f3993, plain, (~ (e3 = op(op(e3, e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3992, f362])).
fof(f3992, plain, (~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3991, f337])).
fof(f337, plain, (e3 = op(e1, e6)), inference(cnf_transformation, [], [f2])).
fof(f3991, plain, (~ (e3 = op(e1, e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3990, f358])).
fof(f3990, plain, (~ (e3 = op(op(e4, e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3989, f369])).
fof(f3989, plain, (~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3988, f359])).
fof(f3988, plain, (~ (e4 = op(e5, e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3987, f345])).
fof(f3987, plain, (~ (e4 = op(op(e3, e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3986, f328])).
fof(f3986, plain, (~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3985, f339])).
fof(f3985, plain, (~ (e4 = op(e2, e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3984, f360])).
fof(f3984, plain, (~ (e4 = op(op(e5, e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3983, f335])).
fof(f335, plain, (e5 = op(e1, e4)), inference(cnf_transformation, [], [f2])).
fof(f3983, plain, (~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3982, f333])).
fof(f3982, plain, (~ (e4 = op(e1, e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3981, f326])).
fof(f3981, plain, (~ (e4 = op(op(e0, e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3980, f342])).
fof(f3980, plain, (~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3979, f369])).
fof(f3979, plain, (~ (e4 = op(e6, e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3978, f341])).
fof(f3978, plain, (~ (e4 = op(op(e2, e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3977, f349])).
fof(f3977, plain, (~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3976, f356])).
fof(f356, plain, (e4 = op(e4, e4)), inference(cnf_transformation, [], [f2])).
fof(f3976, plain, (~ (e4 = op(e4, e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3975, f356])).
fof(f3975, plain, (~ (e4 = op(op(e4, e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3974, f356])).
fof(f3974, plain, (~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3973, f329])).
fof(f329, plain, (e4 = op(e0, e5)), inference(cnf_transformation, [], [f2])).
fof(f3973, plain, (~ (e4 = op(e0, e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3972, f371])).
fof(f3972, plain, (~ (e4 = op(op(e6, e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3971, f363])).
fof(f3971, plain, (~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3970, f351])).
fof(f3970, plain, (~ (e4 = op(e3, e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3969, f337])).
fof(f3969, plain, (~ (e4 = op(op(e1, e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3968, f370])).
fof(f3968, plain, (~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3967, f345])).
fof(f3967, plain, (~ (e5 = op(e3, e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3966, f352])).
fof(f3966, plain, (~ (e5 = op(op(e4, e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3965, f329])).
fof(f3965, plain, (~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3964, f353])).
fof(f3964, plain, (~ (e5 = op(e4, e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3963, f339])).
fof(f3963, plain, (~ (e5 = op(op(e2, e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3962, f336])).
fof(f3962, plain, (~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3961, f361])).
fof(f3961, plain, (~ (e5 = op(e5, e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3960, f361])).
fof(f3960, plain, (~ (e5 = op(op(e5, e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3959, f343])).
fof(f3959, plain, (~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3958, f327])).
fof(f3958, plain, (~ (e5 = op(e0, e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3957, f348])).
fof(f3957, plain, (~ (e5 = op(op(e3, e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3956, f350])).
fof(f3956, plain, (~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3955, f335])).
fof(f3955, plain, (~ (e5 = op(e1, e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3954, f370])).
fof(f3954, plain, (~ (e5 = op(op(e6, e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3953, f357])).
fof(f3953, plain, (~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3952, f343])).
fof(f3952, plain, (~ (e5 = op(e2, e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3951, f336])).
fof(f3951, plain, (~ (e5 = op(op(e1, e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3950, f364])).
fof(f3950, plain, (~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3949, f372])).
fof(f3949, plain, (~ (e5 = op(e6, e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3948, f330])).
fof(f330, plain, (e6 = op(e0, e6)), inference(cnf_transformation, [], [f2])).
fof(f3948, plain, (~ (e5 = op(op(e0, e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3947, f371])).
fof(f3947, plain, (~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3946, f366])).
fof(f3946, plain, (~ (e6 = op(e6, e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3945, f366])).
fof(f3945, plain, (~ (e6 = op(op(e6, e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3944, f330])).
fof(f3944, plain, (~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3943, f332])).
fof(f3943, plain, (~ (e6 = op(e1, e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3942, f346])).
fof(f3942, plain, (~ (e6 = op(op(e3, e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3941, f337])).
fof(f3941, plain, (~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3940, f347])).
fof(f3940, plain, (~ (e6 = op(e3, e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3939, f340])).
fof(f3939, plain, (~ (e6 = op(op(e2, e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3938, f344])).
fof(f3938, plain, (~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3937, f341])).
fof(f3937, plain, (~ (e6 = op(e2, e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3936, f355])).
fof(f3936, plain, (~ (e6 = op(op(e4, e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3935, f351])).
fof(f3935, plain, (~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3934, f363])).
fof(f3934, plain, (~ (e6 = op(e5, e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3933, f335])).
fof(f3933, plain, (~ (e6 = op(op(e1, e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3932, f358])).
fof(f3932, plain, (~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3931, f357])).
fof(f3931, plain, (~ (e6 = op(e4, e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3930, f329])).
fof(f3930, plain, (~ (e6 = op(op(e0, e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3929, f365])).
fof(f3929, plain, (~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3928, f330])).
fof(f3928, plain, (~ (e6 = op(e0, e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3927, f365])).
fof(f3927, plain, (~ (e6 = op(op(e5, e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(forward_demodulation, [], [f3926, f372])).
fof(f3926, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0))), inference(subsumption_resolution, [], [f3925, f2001])).
fof(f2001, plain, ~ sP0, inference(subsumption_resolution, [], [f1406, f356])).
fof(f1406, plain, (~ (e4 = op(e4, e4)) | ~ sP0), inference(cnf_transformation, [], [f302])).
fof(f302, plain, ((~ (e6 = op(e6, e6)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(nnf_transformation, [], [f6])).
fof(f6, plain, ((~ (e6 = op(e6, e6)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e0))) | ~ sP0), inference(usedef, [], [e6])).
fof(e6, plain, (sP0 <=> (~ (e6 = op(e6, e6)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f3925, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP0), inference(subsumption_resolution, [], [f3924, f1998])).
fof(f1998, plain, ~ sP1, inference(subsumption_resolution, [], [f1997, f323])).
fof(f323, plain, ~ (e5 = e6), inference(cnf_transformation, [], [f1])).
fof(f1, plain, (~ (e5 = e6) & ~ (e4 = e6) & ~ (e4 = e5) & ~ (e3 = e6) & ~ (e3 = e5) & ~ (e3 = e4) & ~ (e2 = e6) & ~ (e2 = e5) & ~ (e2 = e4) & ~ (e2 = e3) & ~ (e1 = e6) & ~ (e1 = e5) & ~ (e1 = e4) & ~ (e1 = e3) & ~ (e1 = e2) & ~ (e0 = e6) & ~ (e0 = e5) & ~ (e0 = e4) & ~ (e0 = e3) & ~ (e0 = e2) & ~ (e0 = e1)), file('/home/ubuntu/library/tptp/Problems/ALG/ALG209+1.p', ax1)).
fof(f1997, plain, ((e5 = e6) | ~ sP1), inference(forward_demodulation, [], [f1401, f372])).
fof(f1401, plain, ((e6 = op(e6, e6)) | ~ sP1), inference(cnf_transformation, [], [f301])).
fof(f301, plain, (((e6 = op(e6, e6)) & (e5 = op(e5, e5)) & (e4 = op(e4, e4)) & (e3 = op(e3, e3)) & (e2 = op(e2, e2)) & (e1 = op(e1, e1)) & (e0 = op(e0, e0))) | ~ sP1), inference(nnf_transformation, [], [f7])).
fof(f7, plain, (((e6 = op(e6, e6)) & (e5 = op(e5, e5)) & (e4 = op(e4, e4)) & (e3 = op(e3, e3)) & (e2 = op(e2, e2)) & (e1 = op(e1, e1)) & (e0 = op(e0, e0))) | ~ sP1), inference(usedef, [], [e7])).
fof(e7, plain, (sP1 <=> ((e6 = op(e6, e6)) & (e5 = op(e5, e5)) & (e4 = op(e4, e4)) & (e3 = op(e3, e3)) & (e2 = op(e2, e2)) & (e1 = op(e1, e1)) & (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f3924, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP1 | sP0), inference(subsumption_resolution, [], [f3923, f1996])).
fof(f1996, plain, ~ sP2, inference(subsumption_resolution, [], [f1390, f324])).
fof(f1390, plain, (~ (e2 = op(e0, e0)) | ~ sP2), inference(cnf_transformation, [], [f300])).
fof(f300, plain, ((~ (e6 = op(e0, e0)) & ~ (e5 = op(e0, e0)) & ~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP2), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ((~ (e6 = op(e0, e0)) & ~ (e5 = op(e0, e0)) & ~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ~ sP2), inference(usedef, [], [e8])).
fof(e8, plain, (sP2 <=> (~ (e6 = op(e0, e0)) & ~ (e5 = op(e0, e0)) & ~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f3923, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3922, f1991])).
fof(f1991, plain, ~ sP3, inference(subsumption_resolution, [], [f1381, f325])).
fof(f1381, plain, (~ (e0 = op(e0, e1)) | ~ sP3), inference(cnf_transformation, [], [f299])).
fof(f299, plain, ((~ (e6 = op(e0, e1)) & ~ (e5 = op(e0, e1)) & ~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP3), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ((~ (e6 = op(e0, e1)) & ~ (e5 = op(e0, e1)) & ~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | ~ sP3), inference(usedef, [], [e9])).
fof(e9, plain, (sP3 <=> (~ (e6 = op(e0, e1)) & ~ (e5 = op(e0, e1)) & ~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f3922, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3921, f1984])).
fof(f1984, plain, ~ sP4, inference(subsumption_resolution, [], [f1375, f326])).
fof(f1375, plain, (~ (e1 = op(e0, e2)) | ~ sP4), inference(cnf_transformation, [], [f298])).
fof(f298, plain, ((~ (e6 = op(e0, e2)) & ~ (e5 = op(e0, e2)) & ~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP4), inference(nnf_transformation, [], [f10])).
fof(f10, plain, ((~ (e6 = op(e0, e2)) & ~ (e5 = op(e0, e2)) & ~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | ~ sP4), inference(usedef, [], [e10])).
fof(e10, plain, (sP4 <=> (~ (e6 = op(e0, e2)) & ~ (e5 = op(e0, e2)) & ~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f3921, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3920, f1978])).
fof(f1978, plain, ~ sP5, inference(subsumption_resolution, [], [f1372, f327])).
fof(f1372, plain, (~ (e5 = op(e0, e3)) | ~ sP5), inference(cnf_transformation, [], [f297])).
fof(f297, plain, ((~ (e6 = op(e0, e3)) & ~ (e5 = op(e0, e3)) & ~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP5), inference(nnf_transformation, [], [f11])).
fof(f11, plain, ((~ (e6 = op(e0, e3)) & ~ (e5 = op(e0, e3)) & ~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | ~ sP5), inference(usedef, [], [e11])).
fof(e11, plain, (sP5 <=> (~ (e6 = op(e0, e3)) & ~ (e5 = op(e0, e3)) & ~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f3920, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3919, f1976])).
fof(f1976, plain, ~ sP6, inference(subsumption_resolution, [], [f1363, f328])).
fof(f1363, plain, (~ (e3 = op(e0, e4)) | ~ sP6), inference(cnf_transformation, [], [f296])).
fof(f296, plain, ((~ (e6 = op(e0, e4)) & ~ (e5 = op(e0, e4)) & ~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP6), inference(nnf_transformation, [], [f12])).
fof(f12, plain, ((~ (e6 = op(e0, e4)) & ~ (e5 = op(e0, e4)) & ~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | ~ sP6), inference(usedef, [], [e12])).
fof(e12, plain, (sP6 <=> (~ (e6 = op(e0, e4)) & ~ (e5 = op(e0, e4)) & ~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f3919, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3918, f1972])).
fof(f1972, plain, ~ sP7, inference(subsumption_resolution, [], [f1357, f329])).
fof(f1357, plain, (~ (e4 = op(e0, e5)) | ~ sP7), inference(cnf_transformation, [], [f295])).
fof(f295, plain, ((~ (e6 = op(e0, e5)) & ~ (e5 = op(e0, e5)) & ~ (e4 = op(e0, e5)) & ~ (e3 = op(e0, e5)) & ~ (e2 = op(e0, e5)) & ~ (e1 = op(e0, e5)) & ~ (e0 = op(e0, e5))) | ~ sP7), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ((~ (e6 = op(e0, e5)) & ~ (e5 = op(e0, e5)) & ~ (e4 = op(e0, e5)) & ~ (e3 = op(e0, e5)) & ~ (e2 = op(e0, e5)) & ~ (e1 = op(e0, e5)) & ~ (e0 = op(e0, e5))) | ~ sP7), inference(usedef, [], [e13])).
fof(e13, plain, (sP7 <=> (~ (e6 = op(e0, e5)) & ~ (e5 = op(e0, e5)) & ~ (e4 = op(e0, e5)) & ~ (e3 = op(e0, e5)) & ~ (e2 = op(e0, e5)) & ~ (e1 = op(e0, e5)) & ~ (e0 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f3918, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3917, f1969])).
fof(f1969, plain, ~ sP8, inference(subsumption_resolution, [], [f1352, f330])).
fof(f1352, plain, (~ (e6 = op(e0, e6)) | ~ sP8), inference(cnf_transformation, [], [f294])).
fof(f294, plain, ((~ (e6 = op(e0, e6)) & ~ (e5 = op(e0, e6)) & ~ (e4 = op(e0, e6)) & ~ (e3 = op(e0, e6)) & ~ (e2 = op(e0, e6)) & ~ (e1 = op(e0, e6)) & ~ (e0 = op(e0, e6))) | ~ sP8), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ((~ (e6 = op(e0, e6)) & ~ (e5 = op(e0, e6)) & ~ (e4 = op(e0, e6)) & ~ (e3 = op(e0, e6)) & ~ (e2 = op(e0, e6)) & ~ (e1 = op(e0, e6)) & ~ (e0 = op(e0, e6))) | ~ sP8), inference(usedef, [], [e14])).
fof(e14, plain, (sP8 <=> (~ (e6 = op(e0, e6)) & ~ (e5 = op(e0, e6)) & ~ (e4 = op(e0, e6)) & ~ (e3 = op(e0, e6)) & ~ (e2 = op(e0, e6)) & ~ (e1 = op(e0, e6)) & ~ (e0 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f3917, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3916, f1968])).
fof(f1968, plain, ~ sP9, inference(subsumption_resolution, [], [f1339, f331])).
fof(f1339, plain, (~ (e0 = op(e1, e0)) | ~ sP9), inference(cnf_transformation, [], [f293])).
fof(f293, plain, ((~ (e6 = op(e1, e0)) & ~ (e5 = op(e1, e0)) & ~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP9), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ((~ (e6 = op(e1, e0)) & ~ (e5 = op(e1, e0)) & ~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | ~ sP9), inference(usedef, [], [e15])).
fof(e15, plain, (sP9 <=> (~ (e6 = op(e1, e0)) & ~ (e5 = op(e1, e0)) & ~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f3916, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3915, f1961])).
fof(f1961, plain, ~ sP10, inference(subsumption_resolution, [], [f1338, f332])).
fof(f1338, plain, (~ (e6 = op(e1, e1)) | ~ sP10), inference(cnf_transformation, [], [f292])).
fof(f292, plain, ((~ (e6 = op(e1, e1)) & ~ (e5 = op(e1, e1)) & ~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP10), inference(nnf_transformation, [], [f16])).
fof(f16, plain, ((~ (e6 = op(e1, e1)) & ~ (e5 = op(e1, e1)) & ~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | ~ sP10), inference(usedef, [], [e16])).
fof(e16, plain, (sP10 <=> (~ (e6 = op(e1, e1)) & ~ (e5 = op(e1, e1)) & ~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f3915, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3914, f1960])).
fof(f1960, plain, ~ sP11, inference(subsumption_resolution, [], [f1329, f333])).
fof(f1329, plain, (~ (e4 = op(e1, e2)) | ~ sP11), inference(cnf_transformation, [], [f291])).
fof(f291, plain, ((~ (e6 = op(e1, e2)) & ~ (e5 = op(e1, e2)) & ~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP11), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ((~ (e6 = op(e1, e2)) & ~ (e5 = op(e1, e2)) & ~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | ~ sP11), inference(usedef, [], [e17])).
fof(e17, plain, (sP11 <=> (~ (e6 = op(e1, e2)) & ~ (e5 = op(e1, e2)) & ~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f3914, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3913, f1957])).
fof(f1957, plain, ~ sP12, inference(subsumption_resolution, [], [f1319, f334])).
fof(f1319, plain, (~ (e1 = op(e1, e3)) | ~ sP12), inference(cnf_transformation, [], [f290])).
fof(f290, plain, ((~ (e6 = op(e1, e3)) & ~ (e5 = op(e1, e3)) & ~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP12), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ((~ (e6 = op(e1, e3)) & ~ (e5 = op(e1, e3)) & ~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | ~ sP12), inference(usedef, [], [e18])).
fof(e18, plain, (sP12 <=> (~ (e6 = op(e1, e3)) & ~ (e5 = op(e1, e3)) & ~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f3913, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3912, f1951])).
fof(f1951, plain, ~ sP13, inference(subsumption_resolution, [], [f1316, f335])).
fof(f1316, plain, (~ (e5 = op(e1, e4)) | ~ sP13), inference(cnf_transformation, [], [f289])).
fof(f289, plain, ((~ (e6 = op(e1, e4)) & ~ (e5 = op(e1, e4)) & ~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP13), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ((~ (e6 = op(e1, e4)) & ~ (e5 = op(e1, e4)) & ~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | ~ sP13), inference(usedef, [], [e19])).
fof(e19, plain, (sP13 <=> (~ (e6 = op(e1, e4)) & ~ (e5 = op(e1, e4)) & ~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f3912, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3911, f1949])).
fof(f1949, plain, ~ sP14, inference(subsumption_resolution, [], [f1306, f336])).
fof(f1306, plain, (~ (e2 = op(e1, e5)) | ~ sP14), inference(cnf_transformation, [], [f288])).
fof(f288, plain, ((~ (e6 = op(e1, e5)) & ~ (e5 = op(e1, e5)) & ~ (e4 = op(e1, e5)) & ~ (e3 = op(e1, e5)) & ~ (e2 = op(e1, e5)) & ~ (e1 = op(e1, e5)) & ~ (e0 = op(e1, e5))) | ~ sP14), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ((~ (e6 = op(e1, e5)) & ~ (e5 = op(e1, e5)) & ~ (e4 = op(e1, e5)) & ~ (e3 = op(e1, e5)) & ~ (e2 = op(e1, e5)) & ~ (e1 = op(e1, e5)) & ~ (e0 = op(e1, e5))) | ~ sP14), inference(usedef, [], [e20])).
fof(e20, plain, (sP14 <=> (~ (e6 = op(e1, e5)) & ~ (e5 = op(e1, e5)) & ~ (e4 = op(e1, e5)) & ~ (e3 = op(e1, e5)) & ~ (e2 = op(e1, e5)) & ~ (e1 = op(e1, e5)) & ~ (e0 = op(e1, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f3911, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3910, f1944])).
fof(f1944, plain, ~ sP15, inference(subsumption_resolution, [], [f1300, f337])).
fof(f1300, plain, (~ (e3 = op(e1, e6)) | ~ sP15), inference(cnf_transformation, [], [f287])).
fof(f287, plain, ((~ (e6 = op(e1, e6)) & ~ (e5 = op(e1, e6)) & ~ (e4 = op(e1, e6)) & ~ (e3 = op(e1, e6)) & ~ (e2 = op(e1, e6)) & ~ (e1 = op(e1, e6)) & ~ (e0 = op(e1, e6))) | ~ sP15), inference(nnf_transformation, [], [f21])).
fof(f21, plain, ((~ (e6 = op(e1, e6)) & ~ (e5 = op(e1, e6)) & ~ (e4 = op(e1, e6)) & ~ (e3 = op(e1, e6)) & ~ (e2 = op(e1, e6)) & ~ (e1 = op(e1, e6)) & ~ (e0 = op(e1, e6))) | ~ sP15), inference(usedef, [], [e21])).
fof(e21, plain, (sP15 <=> (~ (e6 = op(e1, e6)) & ~ (e5 = op(e1, e6)) & ~ (e4 = op(e1, e6)) & ~ (e3 = op(e1, e6)) & ~ (e2 = op(e1, e6)) & ~ (e1 = op(e1, e6)) & ~ (e0 = op(e1, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f3910, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3909, f1940])).
fof(f1940, plain, ~ sP16, inference(subsumption_resolution, [], [f1291, f338])).
fof(f1291, plain, (~ (e1 = op(e2, e0)) | ~ sP16), inference(cnf_transformation, [], [f286])).
fof(f286, plain, ((~ (e6 = op(e2, e0)) & ~ (e5 = op(e2, e0)) & ~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP16), inference(nnf_transformation, [], [f22])).
fof(f22, plain, ((~ (e6 = op(e2, e0)) & ~ (e5 = op(e2, e0)) & ~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | ~ sP16), inference(usedef, [], [e22])).
fof(e22, plain, (sP16 <=> (~ (e6 = op(e2, e0)) & ~ (e5 = op(e2, e0)) & ~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f3909, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3908, f1934])).
fof(f1934, plain, ~ sP17, inference(subsumption_resolution, [], [f1287, f339])).
fof(f1287, plain, (~ (e4 = op(e2, e1)) | ~ sP17), inference(cnf_transformation, [], [f285])).
fof(f285, plain, ((~ (e6 = op(e2, e1)) & ~ (e5 = op(e2, e1)) & ~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP17), inference(nnf_transformation, [], [f23])).
fof(f23, plain, ((~ (e6 = op(e2, e1)) & ~ (e5 = op(e2, e1)) & ~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | ~ sP17), inference(usedef, [], [e23])).
fof(e23, plain, (sP17 <=> (~ (e6 = op(e2, e1)) & ~ (e5 = op(e2, e1)) & ~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f3908, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3907, f1931])).
fof(f1931, plain, ~ sP18, inference(subsumption_resolution, [], [f1279, f340])).
fof(f1279, plain, (~ (e3 = op(e2, e2)) | ~ sP18), inference(cnf_transformation, [], [f284])).
fof(f284, plain, ((~ (e6 = op(e2, e2)) & ~ (e5 = op(e2, e2)) & ~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP18), inference(nnf_transformation, [], [f24])).
fof(f24, plain, ((~ (e6 = op(e2, e2)) & ~ (e5 = op(e2, e2)) & ~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | ~ sP18), inference(usedef, [], [e24])).
fof(e24, plain, (sP18 <=> (~ (e6 = op(e2, e2)) & ~ (e5 = op(e2, e2)) & ~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f3907, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3906, f1927])).
fof(f1927, plain, ~ sP19, inference(subsumption_resolution, [], [f1275, f341])).
fof(f1275, plain, (~ (e6 = op(e2, e3)) | ~ sP19), inference(cnf_transformation, [], [f283])).
fof(f283, plain, ((~ (e6 = op(e2, e3)) & ~ (e5 = op(e2, e3)) & ~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP19), inference(nnf_transformation, [], [f25])).
fof(f25, plain, ((~ (e6 = op(e2, e3)) & ~ (e5 = op(e2, e3)) & ~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | ~ sP19), inference(usedef, [], [e25])).
fof(e25, plain, (sP19 <=> (~ (e6 = op(e2, e3)) & ~ (e5 = op(e2, e3)) & ~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f3906, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3905, f1926])).
fof(f1926, plain, ~ sP20, inference(subsumption_resolution, [], [f1262, f342])).
fof(f1262, plain, (~ (e0 = op(e2, e4)) | ~ sP20), inference(cnf_transformation, [], [f282])).
fof(f282, plain, ((~ (e6 = op(e2, e4)) & ~ (e5 = op(e2, e4)) & ~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP20), inference(nnf_transformation, [], [f26])).
fof(f26, plain, ((~ (e6 = op(e2, e4)) & ~ (e5 = op(e2, e4)) & ~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | ~ sP20), inference(usedef, [], [e26])).
fof(e26, plain, (sP20 <=> (~ (e6 = op(e2, e4)) & ~ (e5 = op(e2, e4)) & ~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f3905, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3904, f1919])).
fof(f1919, plain, ~ sP21, inference(subsumption_resolution, [], [f1260, f343])).
fof(f1260, plain, (~ (e5 = op(e2, e5)) | ~ sP21), inference(cnf_transformation, [], [f281])).
fof(f281, plain, ((~ (e6 = op(e2, e5)) & ~ (e5 = op(e2, e5)) & ~ (e4 = op(e2, e5)) & ~ (e3 = op(e2, e5)) & ~ (e2 = op(e2, e5)) & ~ (e1 = op(e2, e5)) & ~ (e0 = op(e2, e5))) | ~ sP21), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ((~ (e6 = op(e2, e5)) & ~ (e5 = op(e2, e5)) & ~ (e4 = op(e2, e5)) & ~ (e3 = op(e2, e5)) & ~ (e2 = op(e2, e5)) & ~ (e1 = op(e2, e5)) & ~ (e0 = op(e2, e5))) | ~ sP21), inference(usedef, [], [e27])).
fof(e27, plain, (sP21 <=> (~ (e6 = op(e2, e5)) & ~ (e5 = op(e2, e5)) & ~ (e4 = op(e2, e5)) & ~ (e3 = op(e2, e5)) & ~ (e2 = op(e2, e5)) & ~ (e1 = op(e2, e5)) & ~ (e0 = op(e2, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f3904, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3903, f1917])).
fof(f1917, plain, ~ sP22, inference(subsumption_resolution, [], [f1250, f344])).
fof(f1250, plain, (~ (e2 = op(e2, e6)) | ~ sP22), inference(cnf_transformation, [], [f280])).
fof(f280, plain, ((~ (e6 = op(e2, e6)) & ~ (e5 = op(e2, e6)) & ~ (e4 = op(e2, e6)) & ~ (e3 = op(e2, e6)) & ~ (e2 = op(e2, e6)) & ~ (e1 = op(e2, e6)) & ~ (e0 = op(e2, e6))) | ~ sP22), inference(nnf_transformation, [], [f28])).
fof(f28, plain, ((~ (e6 = op(e2, e6)) & ~ (e5 = op(e2, e6)) & ~ (e4 = op(e2, e6)) & ~ (e3 = op(e2, e6)) & ~ (e2 = op(e2, e6)) & ~ (e1 = op(e2, e6)) & ~ (e0 = op(e2, e6))) | ~ sP22), inference(usedef, [], [e28])).
fof(e28, plain, (sP22 <=> (~ (e6 = op(e2, e6)) & ~ (e5 = op(e2, e6)) & ~ (e4 = op(e2, e6)) & ~ (e3 = op(e2, e6)) & ~ (e2 = op(e2, e6)) & ~ (e1 = op(e2, e6)) & ~ (e0 = op(e2, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f3903, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3902, f1912])).
fof(f1912, plain, ~ sP23, inference(subsumption_resolution, [], [f1246, f345])).
fof(f1246, plain, (~ (e5 = op(e3, e0)) | ~ sP23), inference(cnf_transformation, [], [f279])).
fof(f279, plain, ((~ (e6 = op(e3, e0)) & ~ (e5 = op(e3, e0)) & ~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP23), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ((~ (e6 = op(e3, e0)) & ~ (e5 = op(e3, e0)) & ~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | ~ sP23), inference(usedef, [], [e29])).
fof(e29, plain, (sP23 <=> (~ (e6 = op(e3, e0)) & ~ (e5 = op(e3, e0)) & ~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f3902, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3901, f1910])).
fof(f1910, plain, ~ sP24, inference(subsumption_resolution, [], [f1235, f346])).
fof(f1235, plain, (~ (e1 = op(e3, e1)) | ~ sP24), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ((~ (e6 = op(e3, e1)) & ~ (e5 = op(e3, e1)) & ~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP24), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ((~ (e6 = op(e3, e1)) & ~ (e5 = op(e3, e1)) & ~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | ~ sP24), inference(usedef, [], [e30])).
fof(e30, plain, (sP24 <=> (~ (e6 = op(e3, e1)) & ~ (e5 = op(e3, e1)) & ~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f3901, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3900, f1904])).
fof(f1904, plain, ~ sP25, inference(subsumption_resolution, [], [f1233, f347])).
fof(f1233, plain, (~ (e6 = op(e3, e2)) | ~ sP25), inference(cnf_transformation, [], [f277])).
fof(f277, plain, ((~ (e6 = op(e3, e2)) & ~ (e5 = op(e3, e2)) & ~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP25), inference(nnf_transformation, [], [f31])).
fof(f31, plain, ((~ (e6 = op(e3, e2)) & ~ (e5 = op(e3, e2)) & ~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | ~ sP25), inference(usedef, [], [e31])).
fof(e31, plain, (sP25 <=> (~ (e6 = op(e3, e2)) & ~ (e5 = op(e3, e2)) & ~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f3900, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3899, f1903])).
fof(f1903, plain, ~ sP26, inference(subsumption_resolution, [], [f1220, f348])).
fof(f1220, plain, (~ (e0 = op(e3, e3)) | ~ sP26), inference(cnf_transformation, [], [f276])).
fof(f276, plain, ((~ (e6 = op(e3, e3)) & ~ (e5 = op(e3, e3)) & ~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP26), inference(nnf_transformation, [], [f32])).
fof(f32, plain, ((~ (e6 = op(e3, e3)) & ~ (e5 = op(e3, e3)) & ~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | ~ sP26), inference(usedef, [], [e32])).
fof(e32, plain, (sP26 <=> (~ (e6 = op(e3, e3)) & ~ (e5 = op(e3, e3)) & ~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f3899, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3898, f1896])).
fof(f1896, plain, ~ sP27, inference(subsumption_resolution, [], [f1215, f349])).
fof(f1215, plain, (~ (e2 = op(e3, e4)) | ~ sP27), inference(cnf_transformation, [], [f275])).
fof(f275, plain, ((~ (e6 = op(e3, e4)) & ~ (e5 = op(e3, e4)) & ~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP27), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ((~ (e6 = op(e3, e4)) & ~ (e5 = op(e3, e4)) & ~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | ~ sP27), inference(usedef, [], [e33])).
fof(e33, plain, (sP27 <=> (~ (e6 = op(e3, e4)) & ~ (e5 = op(e3, e4)) & ~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f3898, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3897, f1891])).
fof(f1891, plain, ~ sP28, inference(subsumption_resolution, [], [f1209, f350])).
fof(f1209, plain, (~ (e3 = op(e3, e5)) | ~ sP28), inference(cnf_transformation, [], [f274])).
fof(f274, plain, ((~ (e6 = op(e3, e5)) & ~ (e5 = op(e3, e5)) & ~ (e4 = op(e3, e5)) & ~ (e3 = op(e3, e5)) & ~ (e2 = op(e3, e5)) & ~ (e1 = op(e3, e5)) & ~ (e0 = op(e3, e5))) | ~ sP28), inference(nnf_transformation, [], [f34])).
fof(f34, plain, ((~ (e6 = op(e3, e5)) & ~ (e5 = op(e3, e5)) & ~ (e4 = op(e3, e5)) & ~ (e3 = op(e3, e5)) & ~ (e2 = op(e3, e5)) & ~ (e1 = op(e3, e5)) & ~ (e0 = op(e3, e5))) | ~ sP28), inference(usedef, [], [e34])).
fof(e34, plain, (sP28 <=> (~ (e6 = op(e3, e5)) & ~ (e5 = op(e3, e5)) & ~ (e4 = op(e3, e5)) & ~ (e3 = op(e3, e5)) & ~ (e2 = op(e3, e5)) & ~ (e1 = op(e3, e5)) & ~ (e0 = op(e3, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f3897, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3896, f1887])).
fof(f1887, plain, ~ sP29, inference(subsumption_resolution, [], [f1203, f351])).
fof(f1203, plain, (~ (e4 = op(e3, e6)) | ~ sP29), inference(cnf_transformation, [], [f273])).
fof(f273, plain, ((~ (e6 = op(e3, e6)) & ~ (e5 = op(e3, e6)) & ~ (e4 = op(e3, e6)) & ~ (e3 = op(e3, e6)) & ~ (e2 = op(e3, e6)) & ~ (e1 = op(e3, e6)) & ~ (e0 = op(e3, e6))) | ~ sP29), inference(nnf_transformation, [], [f35])).
fof(f35, plain, ((~ (e6 = op(e3, e6)) & ~ (e5 = op(e3, e6)) & ~ (e4 = op(e3, e6)) & ~ (e3 = op(e3, e6)) & ~ (e2 = op(e3, e6)) & ~ (e1 = op(e3, e6)) & ~ (e0 = op(e3, e6))) | ~ sP29), inference(usedef, [], [e35])).
fof(e35, plain, (sP29 <=> (~ (e6 = op(e3, e6)) & ~ (e5 = op(e3, e6)) & ~ (e4 = op(e3, e6)) & ~ (e3 = op(e3, e6)) & ~ (e2 = op(e3, e6)) & ~ (e1 = op(e3, e6)) & ~ (e0 = op(e3, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f3896, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3895, f1884])).
fof(f1884, plain, ~ sP30, inference(subsumption_resolution, [], [f1195, f352])).
fof(f1195, plain, (~ (e3 = op(e4, e0)) | ~ sP30), inference(cnf_transformation, [], [f272])).
fof(f272, plain, ((~ (e6 = op(e4, e0)) & ~ (e5 = op(e4, e0)) & ~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP30), inference(nnf_transformation, [], [f36])).
fof(f36, plain, ((~ (e6 = op(e4, e0)) & ~ (e5 = op(e4, e0)) & ~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | ~ sP30), inference(usedef, [], [e36])).
fof(e36, plain, (sP30 <=> (~ (e6 = op(e4, e0)) & ~ (e5 = op(e4, e0)) & ~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f3895, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3894, f1880])).
fof(f1880, plain, ~ sP31, inference(subsumption_resolution, [], [f1190, f353])).
fof(f1190, plain, (~ (e5 = op(e4, e1)) | ~ sP31), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ((~ (e6 = op(e4, e1)) & ~ (e5 = op(e4, e1)) & ~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP31), inference(nnf_transformation, [], [f37])).
fof(f37, plain, ((~ (e6 = op(e4, e1)) & ~ (e5 = op(e4, e1)) & ~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | ~ sP31), inference(usedef, [], [e37])).
fof(e37, plain, (sP31 <=> (~ (e6 = op(e4, e1)) & ~ (e5 = op(e4, e1)) & ~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f3894, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3893, f1878])).
fof(f1878, plain, ~ sP32, inference(subsumption_resolution, [], [f1178, f354])).
fof(f1178, plain, (~ (e0 = op(e4, e2)) | ~ sP32), inference(cnf_transformation, [], [f270])).
fof(f270, plain, ((~ (e6 = op(e4, e2)) & ~ (e5 = op(e4, e2)) & ~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP32), inference(nnf_transformation, [], [f38])).
fof(f38, plain, ((~ (e6 = op(e4, e2)) & ~ (e5 = op(e4, e2)) & ~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | ~ sP32), inference(usedef, [], [e38])).
fof(e38, plain, (sP32 <=> (~ (e6 = op(e4, e2)) & ~ (e5 = op(e4, e2)) & ~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f3893, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3892, f1871])).
fof(f1871, plain, ~ sP33, inference(subsumption_resolution, [], [f1173, f355])).
fof(f1173, plain, (~ (e2 = op(e4, e3)) | ~ sP33), inference(cnf_transformation, [], [f269])).
fof(f269, plain, ((~ (e6 = op(e4, e3)) & ~ (e5 = op(e4, e3)) & ~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP33), inference(nnf_transformation, [], [f39])).
fof(f39, plain, ((~ (e6 = op(e4, e3)) & ~ (e5 = op(e4, e3)) & ~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | ~ sP33), inference(usedef, [], [e39])).
fof(e39, plain, (sP33 <=> (~ (e6 = op(e4, e3)) & ~ (e5 = op(e4, e3)) & ~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f3892, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3891, f1866])).
fof(f1866, plain, ~ sP34, inference(subsumption_resolution, [], [f1168, f356])).
fof(f1168, plain, (~ (e4 = op(e4, e4)) | ~ sP34), inference(cnf_transformation, [], [f268])).
fof(f268, plain, ((~ (e6 = op(e4, e4)) & ~ (e5 = op(e4, e4)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP34), inference(nnf_transformation, [], [f40])).
fof(f40, plain, ((~ (e6 = op(e4, e4)) & ~ (e5 = op(e4, e4)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | ~ sP34), inference(usedef, [], [e40])).
fof(e40, plain, (sP34 <=> (~ (e6 = op(e4, e4)) & ~ (e5 = op(e4, e4)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f3891, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3890, f1863])).
fof(f1863, plain, ~ sP35, inference(subsumption_resolution, [], [f1163, f357])).
fof(f1163, plain, (~ (e6 = op(e4, e5)) | ~ sP35), inference(cnf_transformation, [], [f267])).
fof(f267, plain, ((~ (e6 = op(e4, e5)) & ~ (e5 = op(e4, e5)) & ~ (e4 = op(e4, e5)) & ~ (e3 = op(e4, e5)) & ~ (e2 = op(e4, e5)) & ~ (e1 = op(e4, e5)) & ~ (e0 = op(e4, e5))) | ~ sP35), inference(nnf_transformation, [], [f41])).
fof(f41, plain, ((~ (e6 = op(e4, e5)) & ~ (e5 = op(e4, e5)) & ~ (e4 = op(e4, e5)) & ~ (e3 = op(e4, e5)) & ~ (e2 = op(e4, e5)) & ~ (e1 = op(e4, e5)) & ~ (e0 = op(e4, e5))) | ~ sP35), inference(usedef, [], [e41])).
fof(e41, plain, (sP35 <=> (~ (e6 = op(e4, e5)) & ~ (e5 = op(e4, e5)) & ~ (e4 = op(e4, e5)) & ~ (e3 = op(e4, e5)) & ~ (e2 = op(e4, e5)) & ~ (e1 = op(e4, e5)) & ~ (e0 = op(e4, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f3890, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3889, f1862])).
fof(f1862, plain, ~ sP36, inference(subsumption_resolution, [], [f1151, f358])).
fof(f1151, plain, (~ (e1 = op(e4, e6)) | ~ sP36), inference(cnf_transformation, [], [f266])).
fof(f266, plain, ((~ (e6 = op(e4, e6)) & ~ (e5 = op(e4, e6)) & ~ (e4 = op(e4, e6)) & ~ (e3 = op(e4, e6)) & ~ (e2 = op(e4, e6)) & ~ (e1 = op(e4, e6)) & ~ (e0 = op(e4, e6))) | ~ sP36), inference(nnf_transformation, [], [f42])).
fof(f42, plain, ((~ (e6 = op(e4, e6)) & ~ (e5 = op(e4, e6)) & ~ (e4 = op(e4, e6)) & ~ (e3 = op(e4, e6)) & ~ (e2 = op(e4, e6)) & ~ (e1 = op(e4, e6)) & ~ (e0 = op(e4, e6))) | ~ sP36), inference(usedef, [], [e42])).
fof(e42, plain, (sP36 <=> (~ (e6 = op(e4, e6)) & ~ (e5 = op(e4, e6)) & ~ (e4 = op(e4, e6)) & ~ (e3 = op(e4, e6)) & ~ (e2 = op(e4, e6)) & ~ (e1 = op(e4, e6)) & ~ (e0 = op(e4, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f3889, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3888, f1856])).
fof(f1856, plain, ~ sP37, inference(subsumption_resolution, [], [f1147, f359])).
fof(f1147, plain, (~ (e4 = op(e5, e0)) | ~ sP37), inference(cnf_transformation, [], [f265])).
fof(f265, plain, ((~ (e6 = op(e5, e0)) & ~ (e5 = op(e5, e0)) & ~ (e4 = op(e5, e0)) & ~ (e3 = op(e5, e0)) & ~ (e2 = op(e5, e0)) & ~ (e1 = op(e5, e0)) & ~ (e0 = op(e5, e0))) | ~ sP37), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ((~ (e6 = op(e5, e0)) & ~ (e5 = op(e5, e0)) & ~ (e4 = op(e5, e0)) & ~ (e3 = op(e5, e0)) & ~ (e2 = op(e5, e0)) & ~ (e1 = op(e5, e0)) & ~ (e0 = op(e5, e0))) | ~ sP37), inference(usedef, [], [e43])).
fof(e43, plain, (sP37 <=> (~ (e6 = op(e5, e0)) & ~ (e5 = op(e5, e0)) & ~ (e4 = op(e5, e0)) & ~ (e3 = op(e5, e0)) & ~ (e2 = op(e5, e0)) & ~ (e1 = op(e5, e0)) & ~ (e0 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f3888, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3887, f1853])).
fof(f1853, plain, ~ sP38, inference(subsumption_resolution, [], [f1138, f360])).
fof(f1138, plain, (~ (e2 = op(e5, e1)) | ~ sP38), inference(cnf_transformation, [], [f264])).
fof(f264, plain, ((~ (e6 = op(e5, e1)) & ~ (e5 = op(e5, e1)) & ~ (e4 = op(e5, e1)) & ~ (e3 = op(e5, e1)) & ~ (e2 = op(e5, e1)) & ~ (e1 = op(e5, e1)) & ~ (e0 = op(e5, e1))) | ~ sP38), inference(nnf_transformation, [], [f44])).
fof(f44, plain, ((~ (e6 = op(e5, e1)) & ~ (e5 = op(e5, e1)) & ~ (e4 = op(e5, e1)) & ~ (e3 = op(e5, e1)) & ~ (e2 = op(e5, e1)) & ~ (e1 = op(e5, e1)) & ~ (e0 = op(e5, e1))) | ~ sP38), inference(usedef, [], [e44])).
fof(e44, plain, (sP38 <=> (~ (e6 = op(e5, e1)) & ~ (e5 = op(e5, e1)) & ~ (e4 = op(e5, e1)) & ~ (e3 = op(e5, e1)) & ~ (e2 = op(e5, e1)) & ~ (e1 = op(e5, e1)) & ~ (e0 = op(e5, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f3887, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3886, f1848])).
fof(f1848, plain, ~ sP39, inference(subsumption_resolution, [], [f1134, f361])).
fof(f1134, plain, (~ (e5 = op(e5, e2)) | ~ sP39), inference(cnf_transformation, [], [f263])).
fof(f263, plain, ((~ (e6 = op(e5, e2)) & ~ (e5 = op(e5, e2)) & ~ (e4 = op(e5, e2)) & ~ (e3 = op(e5, e2)) & ~ (e2 = op(e5, e2)) & ~ (e1 = op(e5, e2)) & ~ (e0 = op(e5, e2))) | ~ sP39), inference(nnf_transformation, [], [f45])).
fof(f45, plain, ((~ (e6 = op(e5, e2)) & ~ (e5 = op(e5, e2)) & ~ (e4 = op(e5, e2)) & ~ (e3 = op(e5, e2)) & ~ (e2 = op(e5, e2)) & ~ (e1 = op(e5, e2)) & ~ (e0 = op(e5, e2))) | ~ sP39), inference(usedef, [], [e45])).
fof(e45, plain, (sP39 <=> (~ (e6 = op(e5, e2)) & ~ (e5 = op(e5, e2)) & ~ (e4 = op(e5, e2)) & ~ (e3 = op(e5, e2)) & ~ (e2 = op(e5, e2)) & ~ (e1 = op(e5, e2)) & ~ (e0 = op(e5, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f3886, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3885, f1846])).
fof(f1846, plain, ~ sP40, inference(subsumption_resolution, [], [f1125, f362])).
fof(f1125, plain, (~ (e3 = op(e5, e3)) | ~ sP40), inference(cnf_transformation, [], [f262])).
fof(f262, plain, ((~ (e6 = op(e5, e3)) & ~ (e5 = op(e5, e3)) & ~ (e4 = op(e5, e3)) & ~ (e3 = op(e5, e3)) & ~ (e2 = op(e5, e3)) & ~ (e1 = op(e5, e3)) & ~ (e0 = op(e5, e3))) | ~ sP40), inference(nnf_transformation, [], [f46])).
fof(f46, plain, ((~ (e6 = op(e5, e3)) & ~ (e5 = op(e5, e3)) & ~ (e4 = op(e5, e3)) & ~ (e3 = op(e5, e3)) & ~ (e2 = op(e5, e3)) & ~ (e1 = op(e5, e3)) & ~ (e0 = op(e5, e3))) | ~ sP40), inference(usedef, [], [e46])).
fof(e46, plain, (sP40 <=> (~ (e6 = op(e5, e3)) & ~ (e5 = op(e5, e3)) & ~ (e4 = op(e5, e3)) & ~ (e3 = op(e5, e3)) & ~ (e2 = op(e5, e3)) & ~ (e1 = op(e5, e3)) & ~ (e0 = op(e5, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f3885, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3884, f1842])).
fof(f1842, plain, ~ sP41, inference(subsumption_resolution, [], [f1121, f363])).
fof(f1121, plain, (~ (e6 = op(e5, e4)) | ~ sP41), inference(cnf_transformation, [], [f261])).
fof(f261, plain, ((~ (e6 = op(e5, e4)) & ~ (e5 = op(e5, e4)) & ~ (e4 = op(e5, e4)) & ~ (e3 = op(e5, e4)) & ~ (e2 = op(e5, e4)) & ~ (e1 = op(e5, e4)) & ~ (e0 = op(e5, e4))) | ~ sP41), inference(nnf_transformation, [], [f47])).
fof(f47, plain, ((~ (e6 = op(e5, e4)) & ~ (e5 = op(e5, e4)) & ~ (e4 = op(e5, e4)) & ~ (e3 = op(e5, e4)) & ~ (e2 = op(e5, e4)) & ~ (e1 = op(e5, e4)) & ~ (e0 = op(e5, e4))) | ~ sP41), inference(usedef, [], [e47])).
fof(e47, plain, (sP41 <=> (~ (e6 = op(e5, e4)) & ~ (e5 = op(e5, e4)) & ~ (e4 = op(e5, e4)) & ~ (e3 = op(e5, e4)) & ~ (e2 = op(e5, e4)) & ~ (e1 = op(e5, e4)) & ~ (e0 = op(e5, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f3884, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3883, f1841])).
fof(f1841, plain, ~ sP42, inference(subsumption_resolution, [], [f1109, f364])).
fof(f1109, plain, (~ (e1 = op(e5, e5)) | ~ sP42), inference(cnf_transformation, [], [f260])).
fof(f260, plain, ((~ (e6 = op(e5, e5)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e5, e5)) & ~ (e3 = op(e5, e5)) & ~ (e2 = op(e5, e5)) & ~ (e1 = op(e5, e5)) & ~ (e0 = op(e5, e5))) | ~ sP42), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ((~ (e6 = op(e5, e5)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e5, e5)) & ~ (e3 = op(e5, e5)) & ~ (e2 = op(e5, e5)) & ~ (e1 = op(e5, e5)) & ~ (e0 = op(e5, e5))) | ~ sP42), inference(usedef, [], [e48])).
fof(e48, plain, (sP42 <=> (~ (e6 = op(e5, e5)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e5, e5)) & ~ (e3 = op(e5, e5)) & ~ (e2 = op(e5, e5)) & ~ (e1 = op(e5, e5)) & ~ (e0 = op(e5, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f3883, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3882, f1835])).
fof(f1835, plain, ~ sP43, inference(subsumption_resolution, [], [f1101, f365])).
fof(f1101, plain, (~ (e0 = op(e5, e6)) | ~ sP43), inference(cnf_transformation, [], [f259])).
fof(f259, plain, ((~ (e6 = op(e5, e6)) & ~ (e5 = op(e5, e6)) & ~ (e4 = op(e5, e6)) & ~ (e3 = op(e5, e6)) & ~ (e2 = op(e5, e6)) & ~ (e1 = op(e5, e6)) & ~ (e0 = op(e5, e6))) | ~ sP43), inference(nnf_transformation, [], [f49])).
fof(f49, plain, ((~ (e6 = op(e5, e6)) & ~ (e5 = op(e5, e6)) & ~ (e4 = op(e5, e6)) & ~ (e3 = op(e5, e6)) & ~ (e2 = op(e5, e6)) & ~ (e1 = op(e5, e6)) & ~ (e0 = op(e5, e6))) | ~ sP43), inference(usedef, [], [e49])).
fof(e49, plain, (sP43 <=> (~ (e6 = op(e5, e6)) & ~ (e5 = op(e5, e6)) & ~ (e4 = op(e5, e6)) & ~ (e3 = op(e5, e6)) & ~ (e2 = op(e5, e6)) & ~ (e1 = op(e5, e6)) & ~ (e0 = op(e5, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f3882, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3881, f1828])).
fof(f1828, plain, ~ sP44, inference(subsumption_resolution, [], [f1100, f366])).
fof(f1100, plain, (~ (e6 = op(e6, e0)) | ~ sP44), inference(cnf_transformation, [], [f258])).
fof(f258, plain, ((~ (e6 = op(e6, e0)) & ~ (e5 = op(e6, e0)) & ~ (e4 = op(e6, e0)) & ~ (e3 = op(e6, e0)) & ~ (e2 = op(e6, e0)) & ~ (e1 = op(e6, e0)) & ~ (e0 = op(e6, e0))) | ~ sP44), inference(nnf_transformation, [], [f50])).
fof(f50, plain, ((~ (e6 = op(e6, e0)) & ~ (e5 = op(e6, e0)) & ~ (e4 = op(e6, e0)) & ~ (e3 = op(e6, e0)) & ~ (e2 = op(e6, e0)) & ~ (e1 = op(e6, e0)) & ~ (e0 = op(e6, e0))) | ~ sP44), inference(usedef, [], [e50])).
fof(e50, plain, (sP44 <=> (~ (e6 = op(e6, e0)) & ~ (e5 = op(e6, e0)) & ~ (e4 = op(e6, e0)) & ~ (e3 = op(e6, e0)) & ~ (e2 = op(e6, e0)) & ~ (e1 = op(e6, e0)) & ~ (e0 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f3881, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3880, f1827])).
fof(f1827, plain, ~ sP45, inference(subsumption_resolution, [], [f1090, f367])).
fof(f1090, plain, (~ (e3 = op(e6, e1)) | ~ sP45), inference(cnf_transformation, [], [f257])).
fof(f257, plain, ((~ (e6 = op(e6, e1)) & ~ (e5 = op(e6, e1)) & ~ (e4 = op(e6, e1)) & ~ (e3 = op(e6, e1)) & ~ (e2 = op(e6, e1)) & ~ (e1 = op(e6, e1)) & ~ (e0 = op(e6, e1))) | ~ sP45), inference(nnf_transformation, [], [f51])).
fof(f51, plain, ((~ (e6 = op(e6, e1)) & ~ (e5 = op(e6, e1)) & ~ (e4 = op(e6, e1)) & ~ (e3 = op(e6, e1)) & ~ (e2 = op(e6, e1)) & ~ (e1 = op(e6, e1)) & ~ (e0 = op(e6, e1))) | ~ sP45), inference(usedef, [], [e51])).
fof(e51, plain, (sP45 <=> (~ (e6 = op(e6, e1)) & ~ (e5 = op(e6, e1)) & ~ (e4 = op(e6, e1)) & ~ (e3 = op(e6, e1)) & ~ (e2 = op(e6, e1)) & ~ (e1 = op(e6, e1)) & ~ (e0 = op(e6, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f3880, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3879, f1823])).
fof(f1823, plain, ~ sP46, inference(subsumption_resolution, [], [f1082, f368])).
fof(f1082, plain, (~ (e2 = op(e6, e2)) | ~ sP46), inference(cnf_transformation, [], [f256])).
fof(f256, plain, ((~ (e6 = op(e6, e2)) & ~ (e5 = op(e6, e2)) & ~ (e4 = op(e6, e2)) & ~ (e3 = op(e6, e2)) & ~ (e2 = op(e6, e2)) & ~ (e1 = op(e6, e2)) & ~ (e0 = op(e6, e2))) | ~ sP46), inference(nnf_transformation, [], [f52])).
fof(f52, plain, ((~ (e6 = op(e6, e2)) & ~ (e5 = op(e6, e2)) & ~ (e4 = op(e6, e2)) & ~ (e3 = op(e6, e2)) & ~ (e2 = op(e6, e2)) & ~ (e1 = op(e6, e2)) & ~ (e0 = op(e6, e2))) | ~ sP46), inference(usedef, [], [e52])).
fof(e52, plain, (sP46 <=> (~ (e6 = op(e6, e2)) & ~ (e5 = op(e6, e2)) & ~ (e4 = op(e6, e2)) & ~ (e3 = op(e6, e2)) & ~ (e2 = op(e6, e2)) & ~ (e1 = op(e6, e2)) & ~ (e0 = op(e6, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f3879, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3878, f1818])).
fof(f1818, plain, ~ sP47, inference(subsumption_resolution, [], [f1077, f369])).
fof(f1077, plain, (~ (e4 = op(e6, e3)) | ~ sP47), inference(cnf_transformation, [], [f255])).
fof(f255, plain, ((~ (e6 = op(e6, e3)) & ~ (e5 = op(e6, e3)) & ~ (e4 = op(e6, e3)) & ~ (e3 = op(e6, e3)) & ~ (e2 = op(e6, e3)) & ~ (e1 = op(e6, e3)) & ~ (e0 = op(e6, e3))) | ~ sP47), inference(nnf_transformation, [], [f53])).
fof(f53, plain, ((~ (e6 = op(e6, e3)) & ~ (e5 = op(e6, e3)) & ~ (e4 = op(e6, e3)) & ~ (e3 = op(e6, e3)) & ~ (e2 = op(e6, e3)) & ~ (e1 = op(e6, e3)) & ~ (e0 = op(e6, e3))) | ~ sP47), inference(usedef, [], [e53])).
fof(e53, plain, (sP47 <=> (~ (e6 = op(e6, e3)) & ~ (e5 = op(e6, e3)) & ~ (e4 = op(e6, e3)) & ~ (e3 = op(e6, e3)) & ~ (e2 = op(e6, e3)) & ~ (e1 = op(e6, e3)) & ~ (e0 = op(e6, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f3878, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3877, f1815])).
fof(f1815, plain, ~ sP48, inference(subsumption_resolution, [], [f1067, f370])).
fof(f1067, plain, (~ (e1 = op(e6, e4)) | ~ sP48), inference(cnf_transformation, [], [f254])).
fof(f254, plain, ((~ (e6 = op(e6, e4)) & ~ (e5 = op(e6, e4)) & ~ (e4 = op(e6, e4)) & ~ (e3 = op(e6, e4)) & ~ (e2 = op(e6, e4)) & ~ (e1 = op(e6, e4)) & ~ (e0 = op(e6, e4))) | ~ sP48), inference(nnf_transformation, [], [f54])).
fof(f54, plain, ((~ (e6 = op(e6, e4)) & ~ (e5 = op(e6, e4)) & ~ (e4 = op(e6, e4)) & ~ (e3 = op(e6, e4)) & ~ (e2 = op(e6, e4)) & ~ (e1 = op(e6, e4)) & ~ (e0 = op(e6, e4))) | ~ sP48), inference(usedef, [], [e54])).
fof(e54, plain, (sP48 <=> (~ (e6 = op(e6, e4)) & ~ (e5 = op(e6, e4)) & ~ (e4 = op(e6, e4)) & ~ (e3 = op(e6, e4)) & ~ (e2 = op(e6, e4)) & ~ (e1 = op(e6, e4)) & ~ (e0 = op(e6, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f3877, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3876, f1809])).
fof(f1809, plain, ~ sP49, inference(subsumption_resolution, [], [f1059, f371])).
fof(f1059, plain, (~ (e0 = op(e6, e5)) | ~ sP49), inference(cnf_transformation, [], [f253])).
fof(f253, plain, ((~ (e6 = op(e6, e5)) & ~ (e5 = op(e6, e5)) & ~ (e4 = op(e6, e5)) & ~ (e3 = op(e6, e5)) & ~ (e2 = op(e6, e5)) & ~ (e1 = op(e6, e5)) & ~ (e0 = op(e6, e5))) | ~ sP49), inference(nnf_transformation, [], [f55])).
fof(f55, plain, ((~ (e6 = op(e6, e5)) & ~ (e5 = op(e6, e5)) & ~ (e4 = op(e6, e5)) & ~ (e3 = op(e6, e5)) & ~ (e2 = op(e6, e5)) & ~ (e1 = op(e6, e5)) & ~ (e0 = op(e6, e5))) | ~ sP49), inference(usedef, [], [e55])).
fof(e55, plain, (sP49 <=> (~ (e6 = op(e6, e5)) & ~ (e5 = op(e6, e5)) & ~ (e4 = op(e6, e5)) & ~ (e3 = op(e6, e5)) & ~ (e2 = op(e6, e5)) & ~ (e1 = op(e6, e5)) & ~ (e0 = op(e6, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f3876, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3875, f1802])).
fof(f1802, plain, ~ sP50, inference(subsumption_resolution, [], [f1057, f372])).
fof(f1057, plain, (~ (e5 = op(e6, e6)) | ~ sP50), inference(cnf_transformation, [], [f252])).
fof(f252, plain, ((~ (e6 = op(e6, e6)) & ~ (e5 = op(e6, e6)) & ~ (e4 = op(e6, e6)) & ~ (e3 = op(e6, e6)) & ~ (e2 = op(e6, e6)) & ~ (e1 = op(e6, e6)) & ~ (e0 = op(e6, e6))) | ~ sP50), inference(nnf_transformation, [], [f56])).
fof(f56, plain, ((~ (e6 = op(e6, e6)) & ~ (e5 = op(e6, e6)) & ~ (e4 = op(e6, e6)) & ~ (e3 = op(e6, e6)) & ~ (e2 = op(e6, e6)) & ~ (e1 = op(e6, e6)) & ~ (e0 = op(e6, e6))) | ~ sP50), inference(usedef, [], [e56])).
fof(e56, plain, (sP50 <=> (~ (e6 = op(e6, e6)) & ~ (e5 = op(e6, e6)) & ~ (e4 = op(e6, e6)) & ~ (e3 = op(e6, e6)) & ~ (e2 = op(e6, e6)) & ~ (e1 = op(e6, e6)) & ~ (e0 = op(e6, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f3875, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3874, f1800])).
fof(f1800, plain, ~ sP51, inference(subsumption_resolution, [], [f1046, f325])).
fof(f1046, plain, (~ (e0 = op(e0, e1)) | ~ sP51), inference(cnf_transformation, [], [f251])).
fof(f251, plain, ((~ (e0 = op(e0, e6)) & ~ (e0 = op(e0, e5)) & ~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP51), inference(nnf_transformation, [], [f57])).
fof(f57, plain, ((~ (e0 = op(e0, e6)) & ~ (e0 = op(e0, e5)) & ~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | ~ sP51), inference(usedef, [], [e57])).
fof(e57, plain, (sP51 <=> (~ (e0 = op(e0, e6)) & ~ (e0 = op(e0, e5)) & ~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f3874, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3873, f1794])).
fof(f1794, plain, ~ sP52, inference(subsumption_resolution, [], [f1039, f331])).
fof(f1039, plain, (~ (e0 = op(e1, e0)) | ~ sP52), inference(cnf_transformation, [], [f250])).
fof(f250, plain, ((~ (e0 = op(e6, e0)) & ~ (e0 = op(e5, e0)) & ~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP52), inference(nnf_transformation, [], [f58])).
fof(f58, plain, ((~ (e0 = op(e6, e0)) & ~ (e0 = op(e5, e0)) & ~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | ~ sP52), inference(usedef, [], [e58])).
fof(e58, plain, (sP52 <=> (~ (e0 = op(e6, e0)) & ~ (e0 = op(e5, e0)) & ~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f3873, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3872, f1788])).
fof(f1788, plain, ~ sP53, inference(subsumption_resolution, [], [f1033, f326])).
fof(f1033, plain, (~ (e1 = op(e0, e2)) | ~ sP53), inference(cnf_transformation, [], [f249])).
fof(f249, plain, ((~ (e1 = op(e0, e6)) & ~ (e1 = op(e0, e5)) & ~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP53), inference(nnf_transformation, [], [f59])).
fof(f59, plain, ((~ (e1 = op(e0, e6)) & ~ (e1 = op(e0, e5)) & ~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | ~ sP53), inference(usedef, [], [e59])).
fof(e59, plain, (sP53 <=> (~ (e1 = op(e0, e6)) & ~ (e1 = op(e0, e5)) & ~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f3872, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3871, f1783])).
fof(f1783, plain, ~ sP54, inference(subsumption_resolution, [], [f1026, f338])).
fof(f1026, plain, (~ (e1 = op(e2, e0)) | ~ sP54), inference(cnf_transformation, [], [f248])).
fof(f248, plain, ((~ (e1 = op(e6, e0)) & ~ (e1 = op(e5, e0)) & ~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP54), inference(nnf_transformation, [], [f60])).
fof(f60, plain, ((~ (e1 = op(e6, e0)) & ~ (e1 = op(e5, e0)) & ~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | ~ sP54), inference(usedef, [], [e60])).
fof(e60, plain, (sP54 <=> (~ (e1 = op(e6, e0)) & ~ (e1 = op(e5, e0)) & ~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f3871, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3870, f1778])).
fof(f1778, plain, ~ sP55, inference(subsumption_resolution, [], [f1017, f324])).
fof(f1017, plain, (~ (e2 = op(e0, e0)) | ~ sP55), inference(cnf_transformation, [], [f247])).
fof(f247, plain, ((~ (e2 = op(e0, e6)) & ~ (e2 = op(e0, e5)) & ~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP55), inference(nnf_transformation, [], [f61])).
fof(f61, plain, ((~ (e2 = op(e0, e6)) & ~ (e2 = op(e0, e5)) & ~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | ~ sP55), inference(usedef, [], [e61])).
fof(e61, plain, (sP55 <=> (~ (e2 = op(e0, e6)) & ~ (e2 = op(e0, e5)) & ~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f3870, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3869, f1771])).
fof(f1771, plain, ~ sP56, inference(subsumption_resolution, [], [f1010, f324])).
fof(f1010, plain, (~ (e2 = op(e0, e0)) | ~ sP56), inference(cnf_transformation, [], [f246])).
fof(f246, plain, ((~ (e2 = op(e6, e0)) & ~ (e2 = op(e5, e0)) & ~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP56), inference(nnf_transformation, [], [f62])).
fof(f62, plain, ((~ (e2 = op(e6, e0)) & ~ (e2 = op(e5, e0)) & ~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | ~ sP56), inference(usedef, [], [e62])).
fof(e62, plain, (sP56 <=> (~ (e2 = op(e6, e0)) & ~ (e2 = op(e5, e0)) & ~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f3869, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3868, f1764])).
fof(f1764, plain, ~ sP57, inference(subsumption_resolution, [], [f1007, f328])).
fof(f1007, plain, (~ (e3 = op(e0, e4)) | ~ sP57), inference(cnf_transformation, [], [f245])).
fof(f245, plain, ((~ (e3 = op(e0, e6)) & ~ (e3 = op(e0, e5)) & ~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP57), inference(nnf_transformation, [], [f63])).
fof(f63, plain, ((~ (e3 = op(e0, e6)) & ~ (e3 = op(e0, e5)) & ~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | ~ sP57), inference(usedef, [], [e63])).
fof(e63, plain, (sP57 <=> (~ (e3 = op(e0, e6)) & ~ (e3 = op(e0, e5)) & ~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f3868, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3867, f1761])).
fof(f1761, plain, ~ sP58, inference(subsumption_resolution, [], [f1000, f352])).
fof(f1000, plain, (~ (e3 = op(e4, e0)) | ~ sP58), inference(cnf_transformation, [], [f244])).
fof(f244, plain, ((~ (e3 = op(e6, e0)) & ~ (e3 = op(e5, e0)) & ~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP58), inference(nnf_transformation, [], [f64])).
fof(f64, plain, ((~ (e3 = op(e6, e0)) & ~ (e3 = op(e5, e0)) & ~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | ~ sP58), inference(usedef, [], [e64])).
fof(e64, plain, (sP58 <=> (~ (e3 = op(e6, e0)) & ~ (e3 = op(e5, e0)) & ~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f3867, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3866, f1758])).
fof(f1758, plain, ~ sP59, inference(subsumption_resolution, [], [f994, f329])).
fof(f994, plain, (~ (e4 = op(e0, e5)) | ~ sP59), inference(cnf_transformation, [], [f243])).
fof(f243, plain, ((~ (e4 = op(e0, e6)) & ~ (e4 = op(e0, e5)) & ~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP59), inference(nnf_transformation, [], [f65])).
fof(f65, plain, ((~ (e4 = op(e0, e6)) & ~ (e4 = op(e0, e5)) & ~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | ~ sP59), inference(usedef, [], [e65])).
fof(e65, plain, (sP59 <=> (~ (e4 = op(e0, e6)) & ~ (e4 = op(e0, e5)) & ~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f3866, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3865, f1756])).
fof(f1756, plain, ~ sP60, inference(subsumption_resolution, [], [f987, f359])).
fof(f987, plain, (~ (e4 = op(e5, e0)) | ~ sP60), inference(cnf_transformation, [], [f242])).
fof(f242, plain, ((~ (e4 = op(e6, e0)) & ~ (e4 = op(e5, e0)) & ~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP60), inference(nnf_transformation, [], [f66])).
fof(f66, plain, ((~ (e4 = op(e6, e0)) & ~ (e4 = op(e5, e0)) & ~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | ~ sP60), inference(usedef, [], [e66])).
fof(e66, plain, (sP60 <=> (~ (e4 = op(e6, e0)) & ~ (e4 = op(e5, e0)) & ~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f3865, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3864, f1754])).
fof(f1754, plain, ~ sP61, inference(subsumption_resolution, [], [f978, f327])).
fof(f978, plain, (~ (e5 = op(e0, e3)) | ~ sP61), inference(cnf_transformation, [], [f241])).
fof(f241, plain, ((~ (e5 = op(e0, e6)) & ~ (e5 = op(e0, e5)) & ~ (e5 = op(e0, e4)) & ~ (e5 = op(e0, e3)) & ~ (e5 = op(e0, e2)) & ~ (e5 = op(e0, e1)) & ~ (e5 = op(e0, e0))) | ~ sP61), inference(nnf_transformation, [], [f67])).
fof(f67, plain, ((~ (e5 = op(e0, e6)) & ~ (e5 = op(e0, e5)) & ~ (e5 = op(e0, e4)) & ~ (e5 = op(e0, e3)) & ~ (e5 = op(e0, e2)) & ~ (e5 = op(e0, e1)) & ~ (e5 = op(e0, e0))) | ~ sP61), inference(usedef, [], [e67])).
fof(e67, plain, (sP61 <=> (~ (e5 = op(e0, e6)) & ~ (e5 = op(e0, e5)) & ~ (e5 = op(e0, e4)) & ~ (e5 = op(e0, e3)) & ~ (e5 = op(e0, e2)) & ~ (e5 = op(e0, e1)) & ~ (e5 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f3864, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3863, f1750])).
fof(f1750, plain, ~ sP62, inference(subsumption_resolution, [], [f971, f345])).
fof(f971, plain, (~ (e5 = op(e3, e0)) | ~ sP62), inference(cnf_transformation, [], [f240])).
fof(f240, plain, ((~ (e5 = op(e6, e0)) & ~ (e5 = op(e5, e0)) & ~ (e5 = op(e4, e0)) & ~ (e5 = op(e3, e0)) & ~ (e5 = op(e2, e0)) & ~ (e5 = op(e1, e0)) & ~ (e5 = op(e0, e0))) | ~ sP62), inference(nnf_transformation, [], [f68])).
fof(f68, plain, ((~ (e5 = op(e6, e0)) & ~ (e5 = op(e5, e0)) & ~ (e5 = op(e4, e0)) & ~ (e5 = op(e3, e0)) & ~ (e5 = op(e2, e0)) & ~ (e5 = op(e1, e0)) & ~ (e5 = op(e0, e0))) | ~ sP62), inference(usedef, [], [e68])).
fof(e68, plain, (sP62 <=> (~ (e5 = op(e6, e0)) & ~ (e5 = op(e5, e0)) & ~ (e5 = op(e4, e0)) & ~ (e5 = op(e3, e0)) & ~ (e5 = op(e2, e0)) & ~ (e5 = op(e1, e0)) & ~ (e5 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f3863, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3862, f1746])).
fof(f1746, plain, ~ sP63, inference(subsumption_resolution, [], [f967, f330])).
fof(f967, plain, (~ (e6 = op(e0, e6)) | ~ sP63), inference(cnf_transformation, [], [f239])).
fof(f239, plain, ((~ (e6 = op(e0, e6)) & ~ (e6 = op(e0, e5)) & ~ (e6 = op(e0, e4)) & ~ (e6 = op(e0, e3)) & ~ (e6 = op(e0, e2)) & ~ (e6 = op(e0, e1)) & ~ (e6 = op(e0, e0))) | ~ sP63), inference(nnf_transformation, [], [f69])).
fof(f69, plain, ((~ (e6 = op(e0, e6)) & ~ (e6 = op(e0, e5)) & ~ (e6 = op(e0, e4)) & ~ (e6 = op(e0, e3)) & ~ (e6 = op(e0, e2)) & ~ (e6 = op(e0, e1)) & ~ (e6 = op(e0, e0))) | ~ sP63), inference(usedef, [], [e69])).
fof(e69, plain, (sP63 <=> (~ (e6 = op(e0, e6)) & ~ (e6 = op(e0, e5)) & ~ (e6 = op(e0, e4)) & ~ (e6 = op(e0, e3)) & ~ (e6 = op(e0, e2)) & ~ (e6 = op(e0, e1)) & ~ (e6 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f3862, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3861, f1745])).
fof(f1745, plain, ~ sP64, inference(subsumption_resolution, [], [f960, f366])).
fof(f960, plain, (~ (e6 = op(e6, e0)) | ~ sP64), inference(cnf_transformation, [], [f238])).
fof(f238, plain, ((~ (e6 = op(e6, e0)) & ~ (e6 = op(e5, e0)) & ~ (e6 = op(e4, e0)) & ~ (e6 = op(e3, e0)) & ~ (e6 = op(e2, e0)) & ~ (e6 = op(e1, e0)) & ~ (e6 = op(e0, e0))) | ~ sP64), inference(nnf_transformation, [], [f70])).
fof(f70, plain, ((~ (e6 = op(e6, e0)) & ~ (e6 = op(e5, e0)) & ~ (e6 = op(e4, e0)) & ~ (e6 = op(e3, e0)) & ~ (e6 = op(e2, e0)) & ~ (e6 = op(e1, e0)) & ~ (e6 = op(e0, e0))) | ~ sP64), inference(usedef, [], [e70])).
fof(e70, plain, (sP64 <=> (~ (e6 = op(e6, e0)) & ~ (e6 = op(e5, e0)) & ~ (e6 = op(e4, e0)) & ~ (e6 = op(e3, e0)) & ~ (e6 = op(e2, e0)) & ~ (e6 = op(e1, e0)) & ~ (e6 = op(e0, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f3861, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3860, f1744])).
fof(f1744, plain, ~ sP65, inference(subsumption_resolution, [], [f947, f331])).
fof(f947, plain, (~ (e0 = op(e1, e0)) | ~ sP65), inference(cnf_transformation, [], [f237])).
fof(f237, plain, ((~ (e0 = op(e1, e6)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP65), inference(nnf_transformation, [], [f71])).
fof(f71, plain, ((~ (e0 = op(e1, e6)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | ~ sP65), inference(usedef, [], [e71])).
fof(e71, plain, (sP65 <=> (~ (e0 = op(e1, e6)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f3860, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3859, f1737])).
fof(f1737, plain, ~ sP66, inference(subsumption_resolution, [], [f940, f325])).
fof(f940, plain, (~ (e0 = op(e0, e1)) | ~ sP66), inference(cnf_transformation, [], [f236])).
fof(f236, plain, ((~ (e0 = op(e6, e1)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP66), inference(nnf_transformation, [], [f72])).
fof(f72, plain, ((~ (e0 = op(e6, e1)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | ~ sP66), inference(usedef, [], [e72])).
fof(e72, plain, (sP66 <=> (~ (e0 = op(e6, e1)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f3859, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3858, f1730])).
fof(f1730, plain, ~ sP67, inference(subsumption_resolution, [], [f936, f334])).
fof(f936, plain, (~ (e1 = op(e1, e3)) | ~ sP67), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ((~ (e1 = op(e1, e6)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP67), inference(nnf_transformation, [], [f73])).
fof(f73, plain, ((~ (e1 = op(e1, e6)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | ~ sP67), inference(usedef, [], [e73])).
fof(e73, plain, (sP67 <=> (~ (e1 = op(e1, e6)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f3858, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3857, f1726])).
fof(f1726, plain, ~ sP68, inference(subsumption_resolution, [], [f929, f346])).
fof(f929, plain, (~ (e1 = op(e3, e1)) | ~ sP68), inference(cnf_transformation, [], [f234])).
fof(f234, plain, ((~ (e1 = op(e6, e1)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP68), inference(nnf_transformation, [], [f74])).
fof(f74, plain, ((~ (e1 = op(e6, e1)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | ~ sP68), inference(usedef, [], [e74])).
fof(e74, plain, (sP68 <=> (~ (e1 = op(e6, e1)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f3857, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3856, f1722])).
fof(f1722, plain, ~ sP69, inference(subsumption_resolution, [], [f924, f336])).
fof(f924, plain, (~ (e2 = op(e1, e5)) | ~ sP69), inference(cnf_transformation, [], [f233])).
fof(f233, plain, ((~ (e2 = op(e1, e6)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP69), inference(nnf_transformation, [], [f75])).
fof(f75, plain, ((~ (e2 = op(e1, e6)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | ~ sP69), inference(usedef, [], [e75])).
fof(e75, plain, (sP69 <=> (~ (e2 = op(e1, e6)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f3856, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3855, f1720])).
fof(f1720, plain, ~ sP70, inference(subsumption_resolution, [], [f917, f360])).
fof(f917, plain, (~ (e2 = op(e5, e1)) | ~ sP70), inference(cnf_transformation, [], [f232])).
fof(f232, plain, ((~ (e2 = op(e6, e1)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP70), inference(nnf_transformation, [], [f76])).
fof(f76, plain, ((~ (e2 = op(e6, e1)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | ~ sP70), inference(usedef, [], [e76])).
fof(e76, plain, (sP70 <=> (~ (e2 = op(e6, e1)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f3855, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3854, f1718])).
fof(f1718, plain, ~ sP71, inference(subsumption_resolution, [], [f911, f337])).
fof(f911, plain, (~ (e3 = op(e1, e6)) | ~ sP71), inference(cnf_transformation, [], [f231])).
fof(f231, plain, ((~ (e3 = op(e1, e6)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP71), inference(nnf_transformation, [], [f77])).
fof(f77, plain, ((~ (e3 = op(e1, e6)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | ~ sP71), inference(usedef, [], [e77])).
fof(e77, plain, (sP71 <=> (~ (e3 = op(e1, e6)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f3854, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3853, f1717])).
fof(f1717, plain, ~ sP72, inference(subsumption_resolution, [], [f904, f367])).
fof(f904, plain, (~ (e3 = op(e6, e1)) | ~ sP72), inference(cnf_transformation, [], [f230])).
fof(f230, plain, ((~ (e3 = op(e6, e1)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP72), inference(nnf_transformation, [], [f78])).
fof(f78, plain, ((~ (e3 = op(e6, e1)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | ~ sP72), inference(usedef, [], [e78])).
fof(e78, plain, (sP72 <=> (~ (e3 = op(e6, e1)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f3853, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3852, f1716])).
fof(f1716, plain, ~ sP73, inference(subsumption_resolution, [], [f893, f333])).
fof(f893, plain, (~ (e4 = op(e1, e2)) | ~ sP73), inference(cnf_transformation, [], [f229])).
fof(f229, plain, ((~ (e4 = op(e1, e6)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP73), inference(nnf_transformation, [], [f79])).
fof(f79, plain, ((~ (e4 = op(e1, e6)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | ~ sP73), inference(usedef, [], [e79])).
fof(e79, plain, (sP73 <=> (~ (e4 = op(e1, e6)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f3852, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3851, f1711])).
fof(f1711, plain, ~ sP74, inference(subsumption_resolution, [], [f886, f339])).
fof(f886, plain, (~ (e4 = op(e2, e1)) | ~ sP74), inference(cnf_transformation, [], [f228])).
fof(f228, plain, ((~ (e4 = op(e6, e1)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP74), inference(nnf_transformation, [], [f80])).
fof(f80, plain, ((~ (e4 = op(e6, e1)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | ~ sP74), inference(usedef, [], [e80])).
fof(e80, plain, (sP74 <=> (~ (e4 = op(e6, e1)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f3851, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3850, f1706])).
fof(f1706, plain, ~ sP75, inference(subsumption_resolution, [], [f881, f335])).
fof(f881, plain, (~ (e5 = op(e1, e4)) | ~ sP75), inference(cnf_transformation, [], [f227])).
fof(f227, plain, ((~ (e5 = op(e1, e6)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e1, e0))) | ~ sP75), inference(nnf_transformation, [], [f81])).
fof(f81, plain, ((~ (e5 = op(e1, e6)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e1, e0))) | ~ sP75), inference(usedef, [], [e81])).
fof(e81, plain, (sP75 <=> (~ (e5 = op(e1, e6)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f3850, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3849, f1703])).
fof(f1703, plain, ~ sP76, inference(subsumption_resolution, [], [f874, f353])).
fof(f874, plain, (~ (e5 = op(e4, e1)) | ~ sP76), inference(cnf_transformation, [], [f226])).
fof(f226, plain, ((~ (e5 = op(e6, e1)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e0, e1))) | ~ sP76), inference(nnf_transformation, [], [f82])).
fof(f82, plain, ((~ (e5 = op(e6, e1)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e0, e1))) | ~ sP76), inference(usedef, [], [e82])).
fof(e82, plain, (sP76 <=> (~ (e5 = op(e6, e1)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f3849, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3848, f1700])).
fof(f1700, plain, ~ sP77, inference(subsumption_resolution, [], [f864, f332])).
fof(f864, plain, (~ (e6 = op(e1, e1)) | ~ sP77), inference(cnf_transformation, [], [f225])).
fof(f225, plain, ((~ (e6 = op(e1, e6)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e1, e0))) | ~ sP77), inference(nnf_transformation, [], [f83])).
fof(f83, plain, ((~ (e6 = op(e1, e6)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e1, e0))) | ~ sP77), inference(usedef, [], [e83])).
fof(e83, plain, (sP77 <=> (~ (e6 = op(e1, e6)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e1, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f3848, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3847, f1694])).
fof(f1694, plain, ~ sP78, inference(subsumption_resolution, [], [f857, f332])).
fof(f857, plain, (~ (e6 = op(e1, e1)) | ~ sP78), inference(cnf_transformation, [], [f224])).
fof(f224, plain, ((~ (e6 = op(e6, e1)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e0, e1))) | ~ sP78), inference(nnf_transformation, [], [f84])).
fof(f84, plain, ((~ (e6 = op(e6, e1)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e0, e1))) | ~ sP78), inference(usedef, [], [e84])).
fof(e84, plain, (sP78 <=> (~ (e6 = op(e6, e1)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e0, e1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f3847, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3846, f1688])).
fof(f1688, plain, ~ sP79, inference(subsumption_resolution, [], [f853, f342])).
fof(f853, plain, (~ (e0 = op(e2, e4)) | ~ sP79), inference(cnf_transformation, [], [f223])).
fof(f223, plain, ((~ (e0 = op(e2, e6)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP79), inference(nnf_transformation, [], [f85])).
fof(f85, plain, ((~ (e0 = op(e2, e6)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | ~ sP79), inference(usedef, [], [e85])).
fof(e85, plain, (sP79 <=> (~ (e0 = op(e2, e6)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f3846, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3845, f1685])).
fof(f1685, plain, ~ sP80, inference(subsumption_resolution, [], [f846, f354])).
fof(f846, plain, (~ (e0 = op(e4, e2)) | ~ sP80), inference(cnf_transformation, [], [f222])).
fof(f222, plain, ((~ (e0 = op(e6, e2)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP80), inference(nnf_transformation, [], [f86])).
fof(f86, plain, ((~ (e0 = op(e6, e2)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | ~ sP80), inference(usedef, [], [e86])).
fof(e86, plain, (sP80 <=> (~ (e0 = op(e6, e2)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f3845, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3844, f1682])).
fof(f1682, plain, ~ sP81, inference(subsumption_resolution, [], [f835, f338])).
fof(f835, plain, (~ (e1 = op(e2, e0)) | ~ sP81), inference(cnf_transformation, [], [f221])).
fof(f221, plain, ((~ (e1 = op(e2, e6)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP81), inference(nnf_transformation, [], [f87])).
fof(f87, plain, ((~ (e1 = op(e2, e6)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | ~ sP81), inference(usedef, [], [e87])).
fof(e87, plain, (sP81 <=> (~ (e1 = op(e2, e6)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f3844, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3843, f1675])).
fof(f1675, plain, ~ sP82, inference(subsumption_resolution, [], [f828, f326])).
fof(f828, plain, (~ (e1 = op(e0, e2)) | ~ sP82), inference(cnf_transformation, [], [f220])).
fof(f220, plain, ((~ (e1 = op(e6, e2)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP82), inference(nnf_transformation, [], [f88])).
fof(f88, plain, ((~ (e1 = op(e6, e2)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | ~ sP82), inference(usedef, [], [e88])).
fof(e88, plain, (sP82 <=> (~ (e1 = op(e6, e2)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f3843, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3842, f1668])).
fof(f1668, plain, ~ sP83, inference(subsumption_resolution, [], [f827, f344])).
fof(f827, plain, (~ (e2 = op(e2, e6)) | ~ sP83), inference(cnf_transformation, [], [f219])).
fof(f219, plain, ((~ (e2 = op(e2, e6)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP83), inference(nnf_transformation, [], [f89])).
fof(f89, plain, ((~ (e2 = op(e2, e6)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | ~ sP83), inference(usedef, [], [e89])).
fof(e89, plain, (sP83 <=> (~ (e2 = op(e2, e6)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f3842, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3841, f1667])).
fof(f1667, plain, ~ sP84, inference(subsumption_resolution, [], [f820, f368])).
fof(f820, plain, (~ (e2 = op(e6, e2)) | ~ sP84), inference(cnf_transformation, [], [f218])).
fof(f218, plain, ((~ (e2 = op(e6, e2)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP84), inference(nnf_transformation, [], [f90])).
fof(f90, plain, ((~ (e2 = op(e6, e2)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | ~ sP84), inference(usedef, [], [e90])).
fof(e90, plain, (sP84 <=> (~ (e2 = op(e6, e2)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f3841, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3840, f1666])).
fof(f1666, plain, ~ sP85, inference(subsumption_resolution, [], [f809, f340])).
fof(f809, plain, (~ (e3 = op(e2, e2)) | ~ sP85), inference(cnf_transformation, [], [f217])).
fof(f217, plain, ((~ (e3 = op(e2, e6)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP85), inference(nnf_transformation, [], [f91])).
fof(f91, plain, ((~ (e3 = op(e2, e6)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | ~ sP85), inference(usedef, [], [e91])).
fof(e91, plain, (sP85 <=> (~ (e3 = op(e2, e6)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f3840, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3839, f1661])).
fof(f1661, plain, ~ sP86, inference(subsumption_resolution, [], [f802, f340])).
fof(f802, plain, (~ (e3 = op(e2, e2)) | ~ sP86), inference(cnf_transformation, [], [f216])).
fof(f216, plain, ((~ (e3 = op(e6, e2)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP86), inference(nnf_transformation, [], [f92])).
fof(f92, plain, ((~ (e3 = op(e6, e2)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | ~ sP86), inference(usedef, [], [e92])).
fof(e92, plain, (sP86 <=> (~ (e3 = op(e6, e2)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f3839, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3838, f1656])).
fof(f1656, plain, ~ sP87, inference(subsumption_resolution, [], [f794, f339])).
fof(f794, plain, (~ (e4 = op(e2, e1)) | ~ sP87), inference(cnf_transformation, [], [f215])).
fof(f215, plain, ((~ (e4 = op(e2, e6)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP87), inference(nnf_transformation, [], [f93])).
fof(f93, plain, ((~ (e4 = op(e2, e6)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | ~ sP87), inference(usedef, [], [e93])).
fof(e93, plain, (sP87 <=> (~ (e4 = op(e2, e6)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f3838, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3837, f1650])).
fof(f1650, plain, ~ sP88, inference(subsumption_resolution, [], [f787, f333])).
fof(f787, plain, (~ (e4 = op(e1, e2)) | ~ sP88), inference(cnf_transformation, [], [f214])).
fof(f214, plain, ((~ (e4 = op(e6, e2)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP88), inference(nnf_transformation, [], [f94])).
fof(f94, plain, ((~ (e4 = op(e6, e2)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | ~ sP88), inference(usedef, [], [e94])).
fof(e94, plain, (sP88 <=> (~ (e4 = op(e6, e2)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f3837, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3836, f1644])).
fof(f1644, plain, ~ sP89, inference(subsumption_resolution, [], [f784, f343])).
fof(f784, plain, (~ (e5 = op(e2, e5)) | ~ sP89), inference(cnf_transformation, [], [f213])).
fof(f213, plain, ((~ (e5 = op(e2, e6)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e2, e0))) | ~ sP89), inference(nnf_transformation, [], [f95])).
fof(f95, plain, ((~ (e5 = op(e2, e6)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e2, e0))) | ~ sP89), inference(usedef, [], [e95])).
fof(e95, plain, (sP89 <=> (~ (e5 = op(e2, e6)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f3836, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3835, f1642])).
fof(f1642, plain, ~ sP90, inference(subsumption_resolution, [], [f777, f361])).
fof(f777, plain, (~ (e5 = op(e5, e2)) | ~ sP90), inference(cnf_transformation, [], [f212])).
fof(f212, plain, ((~ (e5 = op(e6, e2)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e0, e2))) | ~ sP90), inference(nnf_transformation, [], [f96])).
fof(f96, plain, ((~ (e5 = op(e6, e2)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e0, e2))) | ~ sP90), inference(usedef, [], [e96])).
fof(e96, plain, (sP90 <=> (~ (e5 = op(e6, e2)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f3835, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3834, f1640])).
fof(f1640, plain, ~ sP91, inference(subsumption_resolution, [], [f768, f341])).
fof(f768, plain, (~ (e6 = op(e2, e3)) | ~ sP91), inference(cnf_transformation, [], [f211])).
fof(f211, plain, ((~ (e6 = op(e2, e6)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e2, e0))) | ~ sP91), inference(nnf_transformation, [], [f97])).
fof(f97, plain, ((~ (e6 = op(e2, e6)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e2, e0))) | ~ sP91), inference(usedef, [], [e97])).
fof(e97, plain, (sP91 <=> (~ (e6 = op(e2, e6)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e2, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f3834, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3833, f1636])).
fof(f1636, plain, ~ sP92, inference(subsumption_resolution, [], [f761, f347])).
fof(f761, plain, (~ (e6 = op(e3, e2)) | ~ sP92), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ((~ (e6 = op(e6, e2)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e0, e2))) | ~ sP92), inference(nnf_transformation, [], [f98])).
fof(f98, plain, ((~ (e6 = op(e6, e2)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e0, e2))) | ~ sP92), inference(usedef, [], [e98])).
fof(e98, plain, (sP92 <=> (~ (e6 = op(e6, e2)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e0, e2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f3833, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3832, f1632])).
fof(f1632, plain, ~ sP93, inference(subsumption_resolution, [], [f754, f348])).
fof(f754, plain, (~ (e0 = op(e3, e3)) | ~ sP93), inference(cnf_transformation, [], [f209])).
fof(f209, plain, ((~ (e0 = op(e3, e6)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP93), inference(nnf_transformation, [], [f99])).
fof(f99, plain, ((~ (e0 = op(e3, e6)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | ~ sP93), inference(usedef, [], [e99])).
fof(e99, plain, (sP93 <=> (~ (e0 = op(e3, e6)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f3832, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3831, f1628])).
fof(f1628, plain, ~ sP94, inference(subsumption_resolution, [], [f747, f348])).
fof(f747, plain, (~ (e0 = op(e3, e3)) | ~ sP94), inference(cnf_transformation, [], [f208])).
fof(f208, plain, ((~ (e0 = op(e6, e3)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP94), inference(nnf_transformation, [], [f100])).
fof(f100, plain, ((~ (e0 = op(e6, e3)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | ~ sP94), inference(usedef, [], [e100])).
fof(e100, plain, (sP94 <=> (~ (e0 = op(e6, e3)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f3831, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3830, f1624])).
fof(f1624, plain, ~ sP95, inference(subsumption_resolution, [], [f738, f346])).
fof(f738, plain, (~ (e1 = op(e3, e1)) | ~ sP95), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ((~ (e1 = op(e3, e6)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP95), inference(nnf_transformation, [], [f101])).
fof(f101, plain, ((~ (e1 = op(e3, e6)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | ~ sP95), inference(usedef, [], [e101])).
fof(e101, plain, (sP95 <=> (~ (e1 = op(e3, e6)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f3830, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3829, f1618])).
fof(f1618, plain, ~ sP96, inference(subsumption_resolution, [], [f731, f334])).
fof(f731, plain, (~ (e1 = op(e1, e3)) | ~ sP96), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ((~ (e1 = op(e6, e3)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP96), inference(nnf_transformation, [], [f102])).
fof(f102, plain, ((~ (e1 = op(e6, e3)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | ~ sP96), inference(usedef, [], [e102])).
fof(e102, plain, (sP96 <=> (~ (e1 = op(e6, e3)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f3829, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3828, f1612])).
fof(f1612, plain, ~ sP97, inference(subsumption_resolution, [], [f727, f349])).
fof(f727, plain, (~ (e2 = op(e3, e4)) | ~ sP97), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ((~ (e2 = op(e3, e6)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP97), inference(nnf_transformation, [], [f103])).
fof(f103, plain, ((~ (e2 = op(e3, e6)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | ~ sP97), inference(usedef, [], [e103])).
fof(e103, plain, (sP97 <=> (~ (e2 = op(e3, e6)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f3828, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3827, f1609])).
fof(f1609, plain, ~ sP98, inference(subsumption_resolution, [], [f720, f355])).
fof(f720, plain, (~ (e2 = op(e4, e3)) | ~ sP98), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ((~ (e2 = op(e6, e3)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP98), inference(nnf_transformation, [], [f104])).
fof(f104, plain, ((~ (e2 = op(e6, e3)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | ~ sP98), inference(usedef, [], [e104])).
fof(e104, plain, (sP98 <=> (~ (e2 = op(e6, e3)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f3827, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3826, f1606])).
fof(f1606, plain, ~ sP99, inference(subsumption_resolution, [], [f714, f350])).
fof(f714, plain, (~ (e3 = op(e3, e5)) | ~ sP99), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ((~ (e3 = op(e3, e6)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP99), inference(nnf_transformation, [], [f105])).
fof(f105, plain, ((~ (e3 = op(e3, e6)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | ~ sP99), inference(usedef, [], [e105])).
fof(e105, plain, (sP99 <=> (~ (e3 = op(e3, e6)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f3826, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3825, f1604])).
fof(f1604, plain, ~ sP100, inference(subsumption_resolution, [], [f707, f362])).
fof(f707, plain, (~ (e3 = op(e5, e3)) | ~ sP100), inference(cnf_transformation, [], [f202])).
fof(f202, plain, ((~ (e3 = op(e6, e3)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP100), inference(nnf_transformation, [], [f106])).
fof(f106, plain, ((~ (e3 = op(e6, e3)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | ~ sP100), inference(usedef, [], [e106])).
fof(e106, plain, (sP100 <=> (~ (e3 = op(e6, e3)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f3825, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3824, f1602])).
fof(f1602, plain, ~ sP101, inference(subsumption_resolution, [], [f701, f351])).
fof(f701, plain, (~ (e4 = op(e3, e6)) | ~ sP101), inference(cnf_transformation, [], [f201])).
fof(f201, plain, ((~ (e4 = op(e3, e6)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP101), inference(nnf_transformation, [], [f107])).
fof(f107, plain, ((~ (e4 = op(e3, e6)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | ~ sP101), inference(usedef, [], [e107])).
fof(e107, plain, (sP101 <=> (~ (e4 = op(e3, e6)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f3824, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3823, f1601])).
fof(f1601, plain, ~ sP102, inference(subsumption_resolution, [], [f694, f369])).
fof(f694, plain, (~ (e4 = op(e6, e3)) | ~ sP102), inference(cnf_transformation, [], [f200])).
fof(f200, plain, ((~ (e4 = op(e6, e3)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP102), inference(nnf_transformation, [], [f108])).
fof(f108, plain, ((~ (e4 = op(e6, e3)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | ~ sP102), inference(usedef, [], [e108])).
fof(e108, plain, (sP102 <=> (~ (e4 = op(e6, e3)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f3823, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3822, f1600])).
fof(f1600, plain, ~ sP103, inference(subsumption_resolution, [], [f681, f345])).
fof(f681, plain, (~ (e5 = op(e3, e0)) | ~ sP103), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ((~ (e5 = op(e3, e6)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e3, e0))) | ~ sP103), inference(nnf_transformation, [], [f109])).
fof(f109, plain, ((~ (e5 = op(e3, e6)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e3, e0))) | ~ sP103), inference(usedef, [], [e109])).
fof(e109, plain, (sP103 <=> (~ (e5 = op(e3, e6)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f3822, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3821, f1593])).
fof(f1593, plain, ~ sP104, inference(subsumption_resolution, [], [f674, f327])).
fof(f674, plain, (~ (e5 = op(e0, e3)) | ~ sP104), inference(cnf_transformation, [], [f198])).
fof(f198, plain, ((~ (e5 = op(e6, e3)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e0, e3))) | ~ sP104), inference(nnf_transformation, [], [f110])).
fof(f110, plain, ((~ (e5 = op(e6, e3)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e0, e3))) | ~ sP104), inference(usedef, [], [e110])).
fof(e110, plain, (sP104 <=> (~ (e5 = op(e6, e3)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f3821, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3820, f1586])).
fof(f1586, plain, ~ sP105, inference(subsumption_resolution, [], [f669, f347])).
fof(f669, plain, (~ (e6 = op(e3, e2)) | ~ sP105), inference(cnf_transformation, [], [f197])).
fof(f197, plain, ((~ (e6 = op(e3, e6)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e3, e0))) | ~ sP105), inference(nnf_transformation, [], [f111])).
fof(f111, plain, ((~ (e6 = op(e3, e6)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e3, e0))) | ~ sP105), inference(usedef, [], [e111])).
fof(e111, plain, (sP105 <=> (~ (e6 = op(e3, e6)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e3, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f3820, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3819, f1581])).
fof(f1581, plain, ~ sP106, inference(subsumption_resolution, [], [f662, f341])).
fof(f662, plain, (~ (e6 = op(e2, e3)) | ~ sP106), inference(cnf_transformation, [], [f196])).
fof(f196, plain, ((~ (e6 = op(e6, e3)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e0, e3))) | ~ sP106), inference(nnf_transformation, [], [f112])).
fof(f112, plain, ((~ (e6 = op(e6, e3)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e0, e3))) | ~ sP106), inference(usedef, [], [e112])).
fof(e112, plain, (sP106 <=> (~ (e6 = op(e6, e3)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e0, e3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f3819, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3818, f1576])).
fof(f1576, plain, ~ sP107, inference(subsumption_resolution, [], [f655, f354])).
fof(f655, plain, (~ (e0 = op(e4, e2)) | ~ sP107), inference(cnf_transformation, [], [f195])).
fof(f195, plain, ((~ (e0 = op(e4, e6)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP107), inference(nnf_transformation, [], [f113])).
fof(f113, plain, ((~ (e0 = op(e4, e6)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | ~ sP107), inference(usedef, [], [e113])).
fof(e113, plain, (sP107 <=> (~ (e0 = op(e4, e6)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f3818, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3817, f1571])).
fof(f1571, plain, ~ sP108, inference(subsumption_resolution, [], [f648, f342])).
fof(f648, plain, (~ (e0 = op(e2, e4)) | ~ sP108), inference(cnf_transformation, [], [f194])).
fof(f194, plain, ((~ (e0 = op(e6, e4)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP108), inference(nnf_transformation, [], [f114])).
fof(f114, plain, ((~ (e0 = op(e6, e4)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | ~ sP108), inference(usedef, [], [e114])).
fof(e114, plain, (sP108 <=> (~ (e0 = op(e6, e4)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f3817, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3816, f1566])).
fof(f1566, plain, ~ sP109, inference(subsumption_resolution, [], [f645, f358])).
fof(f645, plain, (~ (e1 = op(e4, e6)) | ~ sP109), inference(cnf_transformation, [], [f193])).
fof(f193, plain, ((~ (e1 = op(e4, e6)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP109), inference(nnf_transformation, [], [f115])).
fof(f115, plain, ((~ (e1 = op(e4, e6)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | ~ sP109), inference(usedef, [], [e115])).
fof(e115, plain, (sP109 <=> (~ (e1 = op(e4, e6)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f3816, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3815, f1565])).
fof(f1565, plain, ~ sP110, inference(subsumption_resolution, [], [f638, f370])).
fof(f638, plain, (~ (e1 = op(e6, e4)) | ~ sP110), inference(cnf_transformation, [], [f192])).
fof(f192, plain, ((~ (e1 = op(e6, e4)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP110), inference(nnf_transformation, [], [f116])).
fof(f116, plain, ((~ (e1 = op(e6, e4)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | ~ sP110), inference(usedef, [], [e116])).
fof(e116, plain, (sP110 <=> (~ (e1 = op(e6, e4)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f3815, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3814, f1564])).
fof(f1564, plain, ~ sP111, inference(subsumption_resolution, [], [f628, f355])).
fof(f628, plain, (~ (e2 = op(e4, e3)) | ~ sP111), inference(cnf_transformation, [], [f191])).
fof(f191, plain, ((~ (e2 = op(e4, e6)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP111), inference(nnf_transformation, [], [f117])).
fof(f117, plain, ((~ (e2 = op(e4, e6)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | ~ sP111), inference(usedef, [], [e117])).
fof(e117, plain, (sP111 <=> (~ (e2 = op(e4, e6)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f3814, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3813, f1560])).
fof(f1560, plain, ~ sP112, inference(subsumption_resolution, [], [f621, f349])).
fof(f621, plain, (~ (e2 = op(e3, e4)) | ~ sP112), inference(cnf_transformation, [], [f190])).
fof(f190, plain, ((~ (e2 = op(e6, e4)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP112), inference(nnf_transformation, [], [f118])).
fof(f118, plain, ((~ (e2 = op(e6, e4)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | ~ sP112), inference(usedef, [], [e118])).
fof(e118, plain, (sP112 <=> (~ (e2 = op(e6, e4)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f3813, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3812, f1556])).
fof(f1556, plain, ~ sP113, inference(subsumption_resolution, [], [f611, f352])).
fof(f611, plain, (~ (e3 = op(e4, e0)) | ~ sP113), inference(cnf_transformation, [], [f189])).
fof(f189, plain, ((~ (e3 = op(e4, e6)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP113), inference(nnf_transformation, [], [f119])).
fof(f119, plain, ((~ (e3 = op(e4, e6)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | ~ sP113), inference(usedef, [], [e119])).
fof(e119, plain, (sP113 <=> (~ (e3 = op(e4, e6)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f3812, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3811, f1549])).
fof(f1549, plain, ~ sP114, inference(subsumption_resolution, [], [f604, f328])).
fof(f604, plain, (~ (e3 = op(e0, e4)) | ~ sP114), inference(cnf_transformation, [], [f188])).
fof(f188, plain, ((~ (e3 = op(e6, e4)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP114), inference(nnf_transformation, [], [f120])).
fof(f120, plain, ((~ (e3 = op(e6, e4)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | ~ sP114), inference(usedef, [], [e120])).
fof(e120, plain, (sP114 <=> (~ (e3 = op(e6, e4)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f3811, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3810, f1542])).
fof(f1542, plain, ~ sP115, inference(subsumption_resolution, [], [f601, f356])).
fof(f601, plain, (~ (e4 = op(e4, e4)) | ~ sP115), inference(cnf_transformation, [], [f187])).
fof(f187, plain, ((~ (e4 = op(e4, e6)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP115), inference(nnf_transformation, [], [f121])).
fof(f121, plain, ((~ (e4 = op(e4, e6)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | ~ sP115), inference(usedef, [], [e121])).
fof(e121, plain, (sP115 <=> (~ (e4 = op(e4, e6)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f3810, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3809, f1539])).
fof(f1539, plain, ~ sP116, inference(subsumption_resolution, [], [f594, f356])).
fof(f594, plain, (~ (e4 = op(e4, e4)) | ~ sP116), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ((~ (e4 = op(e6, e4)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | ~ sP116), inference(nnf_transformation, [], [f122])).
fof(f122, plain, ((~ (e4 = op(e6, e4)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | ~ sP116), inference(usedef, [], [e122])).
fof(e122, plain, (sP116 <=> (~ (e4 = op(e6, e4)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f3809, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3808, f1536])).
fof(f1536, plain, ~ sP117, inference(subsumption_resolution, [], [f584, f353])).
fof(f584, plain, (~ (e5 = op(e4, e1)) | ~ sP117), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ (e5 = op(e4, e6)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e4, e0))) | ~ sP117), inference(nnf_transformation, [], [f123])).
fof(f123, plain, ((~ (e5 = op(e4, e6)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e4, e0))) | ~ sP117), inference(usedef, [], [e123])).
fof(e123, plain, (sP117 <=> (~ (e5 = op(e4, e6)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f3808, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3807, f1530])).
fof(f1530, plain, ~ sP118, inference(subsumption_resolution, [], [f577, f335])).
fof(f577, plain, (~ (e5 = op(e1, e4)) | ~ sP118), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ((~ (e5 = op(e6, e4)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e0, e4))) | ~ sP118), inference(nnf_transformation, [], [f124])).
fof(f124, plain, ((~ (e5 = op(e6, e4)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e0, e4))) | ~ sP118), inference(usedef, [], [e124])).
fof(e124, plain, (sP118 <=> (~ (e5 = op(e6, e4)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f3807, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3806, f1524])).
fof(f1524, plain, ~ sP119, inference(subsumption_resolution, [], [f574, f357])).
fof(f574, plain, (~ (e6 = op(e4, e5)) | ~ sP119), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ((~ (e6 = op(e4, e6)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e4, e0))) | ~ sP119), inference(nnf_transformation, [], [f125])).
fof(f125, plain, ((~ (e6 = op(e4, e6)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e4, e0))) | ~ sP119), inference(usedef, [], [e125])).
fof(e125, plain, (sP119 <=> (~ (e6 = op(e4, e6)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e4, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f3806, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3805, f1522])).
fof(f1522, plain, ~ sP120, inference(subsumption_resolution, [], [f567, f363])).
fof(f567, plain, (~ (e6 = op(e5, e4)) | ~ sP120), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((~ (e6 = op(e6, e4)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e0, e4))) | ~ sP120), inference(nnf_transformation, [], [f126])).
fof(f126, plain, ((~ (e6 = op(e6, e4)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e0, e4))) | ~ sP120), inference(usedef, [], [e126])).
fof(e126, plain, (sP120 <=> (~ (e6 = op(e6, e4)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e0, e4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f3805, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3804, f1520])).
fof(f1520, plain, ~ sP121, inference(subsumption_resolution, [], [f561, f365])).
fof(f561, plain, (~ (e0 = op(e5, e6)) | ~ sP121), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ((~ (e0 = op(e5, e6)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e5, e0))) | ~ sP121), inference(nnf_transformation, [], [f127])).
fof(f127, plain, ((~ (e0 = op(e5, e6)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e5, e0))) | ~ sP121), inference(usedef, [], [e127])).
fof(e127, plain, (sP121 <=> (~ (e0 = op(e5, e6)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f3804, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3803, f1519])).
fof(f1519, plain, ~ sP122, inference(subsumption_resolution, [], [f554, f371])).
fof(f554, plain, (~ (e0 = op(e6, e5)) | ~ sP122), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((~ (e0 = op(e6, e5)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e0, e5))) | ~ sP122), inference(nnf_transformation, [], [f128])).
fof(f128, plain, ((~ (e0 = op(e6, e5)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e0, e5))) | ~ sP122), inference(usedef, [], [e128])).
fof(e128, plain, (sP122 <=> (~ (e0 = op(e6, e5)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f3803, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3802, f1518])).
fof(f1518, plain, ~ sP123, inference(subsumption_resolution, [], [f546, f364])).
fof(f546, plain, (~ (e1 = op(e5, e5)) | ~ sP123), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((~ (e1 = op(e5, e6)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e5, e0))) | ~ sP123), inference(nnf_transformation, [], [f129])).
fof(f129, plain, ((~ (e1 = op(e5, e6)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e5, e0))) | ~ sP123), inference(usedef, [], [e129])).
fof(e129, plain, (sP123 <=> (~ (e1 = op(e5, e6)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f3802, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3801, f1516])).
fof(f1516, plain, ~ sP124, inference(subsumption_resolution, [], [f539, f364])).
fof(f539, plain, (~ (e1 = op(e5, e5)) | ~ sP124), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ((~ (e1 = op(e6, e5)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e0, e5))) | ~ sP124), inference(nnf_transformation, [], [f130])).
fof(f130, plain, ((~ (e1 = op(e6, e5)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e0, e5))) | ~ sP124), inference(usedef, [], [e130])).
fof(e130, plain, (sP124 <=> (~ (e1 = op(e6, e5)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f3801, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3800, f1514])).
fof(f1514, plain, ~ sP125, inference(subsumption_resolution, [], [f528, f360])).
fof(f528, plain, (~ (e2 = op(e5, e1)) | ~ sP125), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ((~ (e2 = op(e5, e6)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e5, e0))) | ~ sP125), inference(nnf_transformation, [], [f131])).
fof(f131, plain, ((~ (e2 = op(e5, e6)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e5, e0))) | ~ sP125), inference(usedef, [], [e131])).
fof(e131, plain, (sP125 <=> (~ (e2 = op(e5, e6)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f3800, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3799, f1508])).
fof(f1508, plain, ~ sP126, inference(subsumption_resolution, [], [f521, f336])).
fof(f521, plain, (~ (e2 = op(e1, e5)) | ~ sP126), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ((~ (e2 = op(e6, e5)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e0, e5))) | ~ sP126), inference(nnf_transformation, [], [f132])).
fof(f132, plain, ((~ (e2 = op(e6, e5)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e0, e5))) | ~ sP126), inference(usedef, [], [e132])).
fof(e132, plain, (sP126 <=> (~ (e2 = op(e6, e5)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f3799, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3798, f1502])).
fof(f1502, plain, ~ sP127, inference(subsumption_resolution, [], [f516, f362])).
fof(f516, plain, (~ (e3 = op(e5, e3)) | ~ sP127), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ((~ (e3 = op(e5, e6)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e5, e0))) | ~ sP127), inference(nnf_transformation, [], [f133])).
fof(f133, plain, ((~ (e3 = op(e5, e6)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e5, e0))) | ~ sP127), inference(usedef, [], [e133])).
fof(e133, plain, (sP127 <=> (~ (e3 = op(e5, e6)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f3798, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3797, f1498])).
fof(f1498, plain, ~ sP128, inference(subsumption_resolution, [], [f509, f350])).
fof(f509, plain, (~ (e3 = op(e3, e5)) | ~ sP128), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ((~ (e3 = op(e6, e5)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e0, e5))) | ~ sP128), inference(nnf_transformation, [], [f134])).
fof(f134, plain, ((~ (e3 = op(e6, e5)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e0, e5))) | ~ sP128), inference(usedef, [], [e134])).
fof(e134, plain, (sP128 <=> (~ (e3 = op(e6, e5)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f3797, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3796, f1494])).
fof(f1494, plain, ~ sP129, inference(subsumption_resolution, [], [f499, f359])).
fof(f499, plain, (~ (e4 = op(e5, e0)) | ~ sP129), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ((~ (e4 = op(e5, e6)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e5, e0))) | ~ sP129), inference(nnf_transformation, [], [f135])).
fof(f135, plain, ((~ (e4 = op(e5, e6)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e5, e0))) | ~ sP129), inference(usedef, [], [e135])).
fof(e135, plain, (sP129 <=> (~ (e4 = op(e5, e6)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f3796, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3795, f1487])).
fof(f1487, plain, ~ sP130, inference(subsumption_resolution, [], [f492, f329])).
fof(f492, plain, (~ (e4 = op(e0, e5)) | ~ sP130), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ((~ (e4 = op(e6, e5)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e0, e5))) | ~ sP130), inference(nnf_transformation, [], [f136])).
fof(f136, plain, ((~ (e4 = op(e6, e5)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e0, e5))) | ~ sP130), inference(usedef, [], [e136])).
fof(e136, plain, (sP130 <=> (~ (e4 = op(e6, e5)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP130])])).
fof(f3795, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3794, f1480])).
fof(f1480, plain, ~ sP131, inference(subsumption_resolution, [], [f487, f361])).
fof(f487, plain, (~ (e5 = op(e5, e2)) | ~ sP131), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ((~ (e5 = op(e5, e6)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e5, e0))) | ~ sP131), inference(nnf_transformation, [], [f137])).
fof(f137, plain, ((~ (e5 = op(e5, e6)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e5, e0))) | ~ sP131), inference(usedef, [], [e137])).
fof(e137, plain, (sP131 <=> (~ (e5 = op(e5, e6)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP131])])).
fof(f3794, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3793, f1475])).
fof(f1475, plain, ~ sP132, inference(subsumption_resolution, [], [f480, f343])).
fof(f480, plain, (~ (e5 = op(e2, e5)) | ~ sP132), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ((~ (e5 = op(e6, e5)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e0, e5))) | ~ sP132), inference(nnf_transformation, [], [f138])).
fof(f138, plain, ((~ (e5 = op(e6, e5)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e0, e5))) | ~ sP132), inference(usedef, [], [e138])).
fof(e138, plain, (sP132 <=> (~ (e5 = op(e6, e5)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP132])])).
fof(f3793, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3792, f1470])).
fof(f1470, plain, ~ sP133, inference(subsumption_resolution, [], [f475, f363])).
fof(f475, plain, (~ (e6 = op(e5, e4)) | ~ sP133), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ((~ (e6 = op(e5, e6)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e5, e0))) | ~ sP133), inference(nnf_transformation, [], [f139])).
fof(f139, plain, ((~ (e6 = op(e5, e6)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e5, e0))) | ~ sP133), inference(usedef, [], [e139])).
fof(e139, plain, (sP133 <=> (~ (e6 = op(e5, e6)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e5, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP133])])).
fof(f3792, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3791, f1467])).
fof(f1467, plain, ~ sP134, inference(subsumption_resolution, [], [f468, f357])).
fof(f468, plain, (~ (e6 = op(e4, e5)) | ~ sP134), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ((~ (e6 = op(e6, e5)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e0, e5))) | ~ sP134), inference(nnf_transformation, [], [f140])).
fof(f140, plain, ((~ (e6 = op(e6, e5)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e0, e5))) | ~ sP134), inference(usedef, [], [e140])).
fof(e140, plain, (sP134 <=> (~ (e6 = op(e6, e5)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e0, e5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP134])])).
fof(f3791, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3790, f1464])).
fof(f1464, plain, ~ sP135, inference(subsumption_resolution, [], [f462, f371])).
fof(f462, plain, (~ (e0 = op(e6, e5)) | ~ sP135), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ((~ (e0 = op(e6, e6)) & ~ (e0 = op(e6, e5)) & ~ (e0 = op(e6, e4)) & ~ (e0 = op(e6, e3)) & ~ (e0 = op(e6, e2)) & ~ (e0 = op(e6, e1)) & ~ (e0 = op(e6, e0))) | ~ sP135), inference(nnf_transformation, [], [f141])).
fof(f141, plain, ((~ (e0 = op(e6, e6)) & ~ (e0 = op(e6, e5)) & ~ (e0 = op(e6, e4)) & ~ (e0 = op(e6, e3)) & ~ (e0 = op(e6, e2)) & ~ (e0 = op(e6, e1)) & ~ (e0 = op(e6, e0))) | ~ sP135), inference(usedef, [], [e141])).
fof(e141, plain, (sP135 <=> (~ (e0 = op(e6, e6)) & ~ (e0 = op(e6, e5)) & ~ (e0 = op(e6, e4)) & ~ (e0 = op(e6, e3)) & ~ (e0 = op(e6, e2)) & ~ (e0 = op(e6, e1)) & ~ (e0 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP135])])).
fof(f3790, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3789, f1462])).
fof(f1462, plain, ~ sP136, inference(subsumption_resolution, [], [f455, f365])).
fof(f455, plain, (~ (e0 = op(e5, e6)) | ~ sP136), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ((~ (e0 = op(e6, e6)) & ~ (e0 = op(e5, e6)) & ~ (e0 = op(e4, e6)) & ~ (e0 = op(e3, e6)) & ~ (e0 = op(e2, e6)) & ~ (e0 = op(e1, e6)) & ~ (e0 = op(e0, e6))) | ~ sP136), inference(nnf_transformation, [], [f142])).
fof(f142, plain, ((~ (e0 = op(e6, e6)) & ~ (e0 = op(e5, e6)) & ~ (e0 = op(e4, e6)) & ~ (e0 = op(e3, e6)) & ~ (e0 = op(e2, e6)) & ~ (e0 = op(e1, e6)) & ~ (e0 = op(e0, e6))) | ~ sP136), inference(usedef, [], [e142])).
fof(e142, plain, (sP136 <=> (~ (e0 = op(e6, e6)) & ~ (e0 = op(e5, e6)) & ~ (e0 = op(e4, e6)) & ~ (e0 = op(e3, e6)) & ~ (e0 = op(e2, e6)) & ~ (e0 = op(e1, e6)) & ~ (e0 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP136])])).
fof(f3789, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3788, f1460])).
fof(f1460, plain, ~ sP137, inference(subsumption_resolution, [], [f447, f370])).
fof(f447, plain, (~ (e1 = op(e6, e4)) | ~ sP137), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ((~ (e1 = op(e6, e6)) & ~ (e1 = op(e6, e5)) & ~ (e1 = op(e6, e4)) & ~ (e1 = op(e6, e3)) & ~ (e1 = op(e6, e2)) & ~ (e1 = op(e6, e1)) & ~ (e1 = op(e6, e0))) | ~ sP137), inference(nnf_transformation, [], [f143])).
fof(f143, plain, ((~ (e1 = op(e6, e6)) & ~ (e1 = op(e6, e5)) & ~ (e1 = op(e6, e4)) & ~ (e1 = op(e6, e3)) & ~ (e1 = op(e6, e2)) & ~ (e1 = op(e6, e1)) & ~ (e1 = op(e6, e0))) | ~ sP137), inference(usedef, [], [e143])).
fof(e143, plain, (sP137 <=> (~ (e1 = op(e6, e6)) & ~ (e1 = op(e6, e5)) & ~ (e1 = op(e6, e4)) & ~ (e1 = op(e6, e3)) & ~ (e1 = op(e6, e2)) & ~ (e1 = op(e6, e1)) & ~ (e1 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP137])])).
fof(f3788, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3787, f1457])).
fof(f1457, plain, ~ sP138, inference(subsumption_resolution, [], [f440, f358])).
fof(f440, plain, (~ (e1 = op(e4, e6)) | ~ sP138), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ((~ (e1 = op(e6, e6)) & ~ (e1 = op(e5, e6)) & ~ (e1 = op(e4, e6)) & ~ (e1 = op(e3, e6)) & ~ (e1 = op(e2, e6)) & ~ (e1 = op(e1, e6)) & ~ (e1 = op(e0, e6))) | ~ sP138), inference(nnf_transformation, [], [f144])).
fof(f144, plain, ((~ (e1 = op(e6, e6)) & ~ (e1 = op(e5, e6)) & ~ (e1 = op(e4, e6)) & ~ (e1 = op(e3, e6)) & ~ (e1 = op(e2, e6)) & ~ (e1 = op(e1, e6)) & ~ (e1 = op(e0, e6))) | ~ sP138), inference(usedef, [], [e144])).
fof(e144, plain, (sP138 <=> (~ (e1 = op(e6, e6)) & ~ (e1 = op(e5, e6)) & ~ (e1 = op(e4, e6)) & ~ (e1 = op(e3, e6)) & ~ (e1 = op(e2, e6)) & ~ (e1 = op(e1, e6)) & ~ (e1 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP138])])).
fof(f3787, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3786, f1454])).
fof(f1454, plain, ~ sP139, inference(subsumption_resolution, [], [f431, f368])).
fof(f431, plain, (~ (e2 = op(e6, e2)) | ~ sP139), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ((~ (e2 = op(e6, e6)) & ~ (e2 = op(e6, e5)) & ~ (e2 = op(e6, e4)) & ~ (e2 = op(e6, e3)) & ~ (e2 = op(e6, e2)) & ~ (e2 = op(e6, e1)) & ~ (e2 = op(e6, e0))) | ~ sP139), inference(nnf_transformation, [], [f145])).
fof(f145, plain, ((~ (e2 = op(e6, e6)) & ~ (e2 = op(e6, e5)) & ~ (e2 = op(e6, e4)) & ~ (e2 = op(e6, e3)) & ~ (e2 = op(e6, e2)) & ~ (e2 = op(e6, e1)) & ~ (e2 = op(e6, e0))) | ~ sP139), inference(usedef, [], [e145])).
fof(e145, plain, (sP139 <=> (~ (e2 = op(e6, e6)) & ~ (e2 = op(e6, e5)) & ~ (e2 = op(e6, e4)) & ~ (e2 = op(e6, e3)) & ~ (e2 = op(e6, e2)) & ~ (e2 = op(e6, e1)) & ~ (e2 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP139])])).
fof(f3786, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3785, f1449])).
fof(f1449, plain, ~ sP140, inference(subsumption_resolution, [], [f424, f344])).
fof(f424, plain, (~ (e2 = op(e2, e6)) | ~ sP140), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ((~ (e2 = op(e6, e6)) & ~ (e2 = op(e5, e6)) & ~ (e2 = op(e4, e6)) & ~ (e2 = op(e3, e6)) & ~ (e2 = op(e2, e6)) & ~ (e2 = op(e1, e6)) & ~ (e2 = op(e0, e6))) | ~ sP140), inference(nnf_transformation, [], [f146])).
fof(f146, plain, ((~ (e2 = op(e6, e6)) & ~ (e2 = op(e5, e6)) & ~ (e2 = op(e4, e6)) & ~ (e2 = op(e3, e6)) & ~ (e2 = op(e2, e6)) & ~ (e2 = op(e1, e6)) & ~ (e2 = op(e0, e6))) | ~ sP140), inference(usedef, [], [e146])).
fof(e146, plain, (sP140 <=> (~ (e2 = op(e6, e6)) & ~ (e2 = op(e5, e6)) & ~ (e2 = op(e4, e6)) & ~ (e2 = op(e3, e6)) & ~ (e2 = op(e2, e6)) & ~ (e2 = op(e1, e6)) & ~ (e2 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP140])])).
fof(f3785, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3784, f1444])).
fof(f1444, plain, ~ sP141, inference(subsumption_resolution, [], [f416, f367])).
fof(f416, plain, (~ (e3 = op(e6, e1)) | ~ sP141), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ((~ (e3 = op(e6, e6)) & ~ (e3 = op(e6, e5)) & ~ (e3 = op(e6, e4)) & ~ (e3 = op(e6, e3)) & ~ (e3 = op(e6, e2)) & ~ (e3 = op(e6, e1)) & ~ (e3 = op(e6, e0))) | ~ sP141), inference(nnf_transformation, [], [f147])).
fof(f147, plain, ((~ (e3 = op(e6, e6)) & ~ (e3 = op(e6, e5)) & ~ (e3 = op(e6, e4)) & ~ (e3 = op(e6, e3)) & ~ (e3 = op(e6, e2)) & ~ (e3 = op(e6, e1)) & ~ (e3 = op(e6, e0))) | ~ sP141), inference(usedef, [], [e147])).
fof(e147, plain, (sP141 <=> (~ (e3 = op(e6, e6)) & ~ (e3 = op(e6, e5)) & ~ (e3 = op(e6, e4)) & ~ (e3 = op(e6, e3)) & ~ (e3 = op(e6, e2)) & ~ (e3 = op(e6, e1)) & ~ (e3 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP141])])).
fof(f3784, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3783, f1438])).
fof(f1438, plain, ~ sP142, inference(subsumption_resolution, [], [f409, f337])).
fof(f409, plain, (~ (e3 = op(e1, e6)) | ~ sP142), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ((~ (e3 = op(e6, e6)) & ~ (e3 = op(e5, e6)) & ~ (e3 = op(e4, e6)) & ~ (e3 = op(e3, e6)) & ~ (e3 = op(e2, e6)) & ~ (e3 = op(e1, e6)) & ~ (e3 = op(e0, e6))) | ~ sP142), inference(nnf_transformation, [], [f148])).
fof(f148, plain, ((~ (e3 = op(e6, e6)) & ~ (e3 = op(e5, e6)) & ~ (e3 = op(e4, e6)) & ~ (e3 = op(e3, e6)) & ~ (e3 = op(e2, e6)) & ~ (e3 = op(e1, e6)) & ~ (e3 = op(e0, e6))) | ~ sP142), inference(usedef, [], [e148])).
fof(e148, plain, (sP142 <=> (~ (e3 = op(e6, e6)) & ~ (e3 = op(e5, e6)) & ~ (e3 = op(e4, e6)) & ~ (e3 = op(e3, e6)) & ~ (e3 = op(e2, e6)) & ~ (e3 = op(e1, e6)) & ~ (e3 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP142])])).
fof(f3783, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3782, f1432])).
fof(f1432, plain, ~ sP143, inference(subsumption_resolution, [], [f404, f369])).
fof(f404, plain, (~ (e4 = op(e6, e3)) | ~ sP143), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ((~ (e4 = op(e6, e6)) & ~ (e4 = op(e6, e5)) & ~ (e4 = op(e6, e4)) & ~ (e4 = op(e6, e3)) & ~ (e4 = op(e6, e2)) & ~ (e4 = op(e6, e1)) & ~ (e4 = op(e6, e0))) | ~ sP143), inference(nnf_transformation, [], [f149])).
fof(f149, plain, ((~ (e4 = op(e6, e6)) & ~ (e4 = op(e6, e5)) & ~ (e4 = op(e6, e4)) & ~ (e4 = op(e6, e3)) & ~ (e4 = op(e6, e2)) & ~ (e4 = op(e6, e1)) & ~ (e4 = op(e6, e0))) | ~ sP143), inference(usedef, [], [e149])).
fof(e149, plain, (sP143 <=> (~ (e4 = op(e6, e6)) & ~ (e4 = op(e6, e5)) & ~ (e4 = op(e6, e4)) & ~ (e4 = op(e6, e3)) & ~ (e4 = op(e6, e2)) & ~ (e4 = op(e6, e1)) & ~ (e4 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP143])])).
fof(f3782, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3781, f1428])).
fof(f1428, plain, ~ sP144, inference(subsumption_resolution, [], [f397, f351])).
fof(f397, plain, (~ (e4 = op(e3, e6)) | ~ sP144), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ((~ (e4 = op(e6, e6)) & ~ (e4 = op(e5, e6)) & ~ (e4 = op(e4, e6)) & ~ (e4 = op(e3, e6)) & ~ (e4 = op(e2, e6)) & ~ (e4 = op(e1, e6)) & ~ (e4 = op(e0, e6))) | ~ sP144), inference(nnf_transformation, [], [f150])).
fof(f150, plain, ((~ (e4 = op(e6, e6)) & ~ (e4 = op(e5, e6)) & ~ (e4 = op(e4, e6)) & ~ (e4 = op(e3, e6)) & ~ (e4 = op(e2, e6)) & ~ (e4 = op(e1, e6)) & ~ (e4 = op(e0, e6))) | ~ sP144), inference(usedef, [], [e150])).
fof(e150, plain, (sP144 <=> (~ (e4 = op(e6, e6)) & ~ (e4 = op(e5, e6)) & ~ (e4 = op(e4, e6)) & ~ (e4 = op(e3, e6)) & ~ (e4 = op(e2, e6)) & ~ (e4 = op(e1, e6)) & ~ (e4 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP144])])).
fof(f3781, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3780, f1424])).
fof(f1424, plain, ~ sP145, inference(subsumption_resolution, [], [f393, f372])).
fof(f393, plain, (~ (e5 = op(e6, e6)) | ~ sP145), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ((~ (e5 = op(e6, e6)) & ~ (e5 = op(e6, e5)) & ~ (e5 = op(e6, e4)) & ~ (e5 = op(e6, e3)) & ~ (e5 = op(e6, e2)) & ~ (e5 = op(e6, e1)) & ~ (e5 = op(e6, e0))) | ~ sP145), inference(nnf_transformation, [], [f151])).
fof(f151, plain, ((~ (e5 = op(e6, e6)) & ~ (e5 = op(e6, e5)) & ~ (e5 = op(e6, e4)) & ~ (e5 = op(e6, e3)) & ~ (e5 = op(e6, e2)) & ~ (e5 = op(e6, e1)) & ~ (e5 = op(e6, e0))) | ~ sP145), inference(usedef, [], [e151])).
fof(e151, plain, (sP145 <=> (~ (e5 = op(e6, e6)) & ~ (e5 = op(e6, e5)) & ~ (e5 = op(e6, e4)) & ~ (e5 = op(e6, e3)) & ~ (e5 = op(e6, e2)) & ~ (e5 = op(e6, e1)) & ~ (e5 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP145])])).
fof(f3780, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP145 | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3779, f1423])).
fof(f1423, plain, ~ sP146, inference(subsumption_resolution, [], [f386, f372])).
fof(f386, plain, (~ (e5 = op(e6, e6)) | ~ sP146), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ((~ (e5 = op(e6, e6)) & ~ (e5 = op(e5, e6)) & ~ (e5 = op(e4, e6)) & ~ (e5 = op(e3, e6)) & ~ (e5 = op(e2, e6)) & ~ (e5 = op(e1, e6)) & ~ (e5 = op(e0, e6))) | ~ sP146), inference(nnf_transformation, [], [f152])).
fof(f152, plain, ((~ (e5 = op(e6, e6)) & ~ (e5 = op(e5, e6)) & ~ (e5 = op(e4, e6)) & ~ (e5 = op(e3, e6)) & ~ (e5 = op(e2, e6)) & ~ (e5 = op(e1, e6)) & ~ (e5 = op(e0, e6))) | ~ sP146), inference(usedef, [], [e152])).
fof(e152, plain, (sP146 <=> (~ (e5 = op(e6, e6)) & ~ (e5 = op(e5, e6)) & ~ (e5 = op(e4, e6)) & ~ (e5 = op(e3, e6)) & ~ (e5 = op(e2, e6)) & ~ (e5 = op(e1, e6)) & ~ (e5 = op(e0, e6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP146])])).
fof(f3779, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f3778, f1422])).
fof(f1422, plain, ~ sP147, inference(subsumption_resolution, [], [f373, f366])).
fof(f373, plain, (~ (e6 = op(e6, e0)) | ~ sP147), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ((~ (e6 = op(e6, e6)) & ~ (e6 = op(e6, e5)) & ~ (e6 = op(e6, e4)) & ~ (e6 = op(e6, e3)) & ~ (e6 = op(e6, e2)) & ~ (e6 = op(e6, e1)) & ~ (e6 = op(e6, e0))) | ~ sP147), inference(nnf_transformation, [], [f153])).
fof(f153, plain, ((~ (e6 = op(e6, e6)) & ~ (e6 = op(e6, e5)) & ~ (e6 = op(e6, e4)) & ~ (e6 = op(e6, e3)) & ~ (e6 = op(e6, e2)) & ~ (e6 = op(e6, e1)) & ~ (e6 = op(e6, e0))) | ~ sP147), inference(usedef, [], [e153])).
fof(e153, plain, (sP147 <=> (~ (e6 = op(e6, e6)) & ~ (e6 = op(e6, e5)) & ~ (e6 = op(e6, e4)) & ~ (e6 = op(e6, e3)) & ~ (e6 = op(e6, e2)) & ~ (e6 = op(e6, e1)) & ~ (e6 = op(e6, e0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP147])])).
fof(f3778, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(subsumption_resolution, [], [f1409, f330])).
fof(f1409, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | ~ (e6 = op(e0, e6)) | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(cnf_transformation, [], [f154])).
fof(f154, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | (~ (e6 = op(e6, e6)) & ~ (e6 = op(e5, e6)) & ~ (e6 = op(e4, e6)) & ~ (e6 = op(e3, e6)) & ~ (e6 = op(e2, e6)) & ~ (e6 = op(e1, e6)) & ~ (e6 = op(e0, e6))) | sP147 | sP146 | sP145 | sP144 | sP143 | sP142 | sP141 | sP140 | sP139 | sP138 | sP137 | sP136 | sP135 | sP134 | sP133 | sP132 | sP131 | sP130 | sP129 | sP128 | sP127 | sP126 | sP125 | sP124 | sP123 | sP122 | sP121 | sP120 | sP119 | sP118 | sP117 | sP116 | sP115 | sP114 | sP113 | sP112 | sP111 | sP110 | sP109 | sP108 | sP107 | sP106 | sP105 | sP104 | sP103 | sP102 | sP101 | sP100 | sP99 | sP98 | sP97 | sP96 | sP95 | sP94 | sP93 | sP92 | sP91 | sP90 | sP89 | sP88 | sP87 | sP86 | sP85 | sP84 | sP83 | sP82 | sP81 | sP80 | sP79 | sP78 | sP77 | sP76 | sP75 | sP74 | sP73 | sP72 | sP71 | sP70 | sP69 | sP68 | sP67 | sP66 | sP65 | sP64 | sP63 | sP62 | sP61 | sP60 | sP59 | sP58 | sP57 | sP56 | sP55 | sP54 | sP53 | sP52 | sP51 | sP50 | sP49 | sP48 | sP47 | sP46 | sP45 | sP44 | sP43 | sP42 | sP41 | sP40 | sP39 | sP38 | sP37 | sP36 | sP35 | sP34 | sP33 | sP32 | sP31 | sP30 | sP29 | sP28 | sP27 | sP26 | sP25 | sP24 | sP23 | sP22 | sP21 | sP20 | sP19 | sP18 | sP17 | sP16 | sP15 | sP14 | sP13 | sP12 | sP11 | sP10 | sP9 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | sP1 | sP0), inference(definition_folding, [], [f5, e153, e152, e151, e150, e149, e148, e147, e146, e145, e144, e143, e142, e141, e140, e139, e138, e137, e136, e135, e134, e133, e132, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120, e119, e118, e117, e116, e115, e114, e113, e112, e111, e110, e109, e108, e107, e106, e105, e104, e103, e102, e101, e100, e99, e98, e97, e96, e95, e94, e93, e92, e91, e90, e89, e88, e87, e86, e85, e84, e83, e82, e81, e80, e79, e78, e77, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7, e6])).
fof(f5, plain, (~ (e6 = op(op(op(e6, e6), e6), e6)) | ~ (e6 = op(op(op(e5, e6), e5), e5)) | ~ (e6 = op(op(op(e4, e6), e4), e4)) | ~ (e6 = op(op(op(e3, e6), e3), e3)) | ~ (e6 = op(op(op(e2, e6), e2), e2)) | ~ (e6 = op(op(op(e1, e6), e1), e1)) | ~ (e6 = op(op(op(e0, e6), e0), e0)) | ~ (e5 = op(op(op(e6, e5), e6), e6)) | ~ (e5 = op(op(op(e5, e5), e5), e5)) | ~ (e5 = op(op(op(e4, e5), e4), e4)) | ~ (e5 = op(op(op(e3, e5), e3), e3)) | ~ (e5 = op(op(op(e2, e5), e2), e2)) | ~ (e5 = op(op(op(e1, e5), e1), e1)) | ~ (e5 = op(op(op(e0, e5), e0), e0)) | ~ (e4 = op(op(op(e6, e4), e6), e6)) | ~ (e4 = op(op(op(e5, e4), e5), e5)) | ~ (e4 = op(op(op(e4, e4), e4), e4)) | ~ (e4 = op(op(op(e3, e4), e3), e3)) | ~ (e4 = op(op(op(e2, e4), e2), e2)) | ~ (e4 = op(op(op(e1, e4), e1), e1)) | ~ (e4 = op(op(op(e0, e4), e0), e0)) | ~ (e3 = op(op(op(e6, e3), e6), e6)) | ~ (e3 = op(op(op(e5, e3), e5), e5)) | ~ (e3 = op(op(op(e4, e3), e4), e4)) | ~ (e3 = op(op(op(e3, e3), e3), e3)) | ~ (e3 = op(op(op(e2, e3), e2), e2)) | ~ (e3 = op(op(op(e1, e3), e1), e1)) | ~ (e3 = op(op(op(e0, e3), e0), e0)) | ~ (e2 = op(op(op(e6, e2), e6), e6)) | ~ (e2 = op(op(op(e5, e2), e5), e5)) | ~ (e2 = op(op(op(e4, e2), e4), e4)) | ~ (e2 = op(op(op(e3, e2), e3), e3)) | ~ (e2 = op(op(op(e2, e2), e2), e2)) | ~ (e2 = op(op(op(e1, e2), e1), e1)) | ~ (e2 = op(op(op(e0, e2), e0), e0)) | ~ (e1 = op(op(op(e6, e1), e6), e6)) | ~ (e1 = op(op(op(e5, e1), e5), e5)) | ~ (e1 = op(op(op(e4, e1), e4), e4)) | ~ (e1 = op(op(op(e3, e1), e3), e3)) | ~ (e1 = op(op(op(e2, e1), e2), e2)) | ~ (e1 = op(op(op(e1, e1), e1), e1)) | ~ (e1 = op(op(op(e0, e1), e0), e0)) | ~ (e0 = op(op(op(e6, e0), e6), e6)) | ~ (e0 = op(op(op(e5, e0), e5), e5)) | ~ (e0 = op(op(op(e4, e0), e4), e4)) | ~ (e0 = op(op(op(e3, e0), e3), e3)) | ~ (e0 = op(op(op(e2, e0), e2), e2)) | ~ (e0 = op(op(op(e1, e0), e1), e1)) | ~ (e0 = op(op(op(e0, e0), e0), e0)) | (~ (e6 = op(e6, e6)) & ~ (e6 = op(e5, e6)) & ~ (e6 = op(e4, e6)) & ~ (e6 = op(e3, e6)) & ~ (e6 = op(e2, e6)) & ~ (e6 = op(e1, e6)) & ~ (e6 = op(e0, e6))) | (~ (e6 = op(e6, e6)) & ~ (e6 = op(e6, e5)) & ~ (e6 = op(e6, e4)) & ~ (e6 = op(e6, e3)) & ~ (e6 = op(e6, e2)) & ~ (e6 = op(e6, e1)) & ~ (e6 = op(e6, e0))) | (~ (e5 = op(e6, e6)) & ~ (e5 = op(e5, e6)) & ~ (e5 = op(e4, e6)) & ~ (e5 = op(e3, e6)) & ~ (e5 = op(e2, e6)) & ~ (e5 = op(e1, e6)) & ~ (e5 = op(e0, e6))) | (~ (e5 = op(e6, e6)) & ~ (e5 = op(e6, e5)) & ~ (e5 = op(e6, e4)) & ~ (e5 = op(e6, e3)) & ~ (e5 = op(e6, e2)) & ~ (e5 = op(e6, e1)) & ~ (e5 = op(e6, e0))) | (~ (e4 = op(e6, e6)) & ~ (e4 = op(e5, e6)) & ~ (e4 = op(e4, e6)) & ~ (e4 = op(e3, e6)) & ~ (e4 = op(e2, e6)) & ~ (e4 = op(e1, e6)) & ~ (e4 = op(e0, e6))) | (~ (e4 = op(e6, e6)) & ~ (e4 = op(e6, e5)) & ~ (e4 = op(e6, e4)) & ~ (e4 = op(e6, e3)) & ~ (e4 = op(e6, e2)) & ~ (e4 = op(e6, e1)) & ~ (e4 = op(e6, e0))) | (~ (e3 = op(e6, e6)) & ~ (e3 = op(e5, e6)) & ~ (e3 = op(e4, e6)) & ~ (e3 = op(e3, e6)) & ~ (e3 = op(e2, e6)) & ~ (e3 = op(e1, e6)) & ~ (e3 = op(e0, e6))) | (~ (e3 = op(e6, e6)) & ~ (e3 = op(e6, e5)) & ~ (e3 = op(e6, e4)) & ~ (e3 = op(e6, e3)) & ~ (e3 = op(e6, e2)) & ~ (e3 = op(e6, e1)) & ~ (e3 = op(e6, e0))) | (~ (e2 = op(e6, e6)) & ~ (e2 = op(e5, e6)) & ~ (e2 = op(e4, e6)) & ~ (e2 = op(e3, e6)) & ~ (e2 = op(e2, e6)) & ~ (e2 = op(e1, e6)) & ~ (e2 = op(e0, e6))) | (~ (e2 = op(e6, e6)) & ~ (e2 = op(e6, e5)) & ~ (e2 = op(e6, e4)) & ~ (e2 = op(e6, e3)) & ~ (e2 = op(e6, e2)) & ~ (e2 = op(e6, e1)) & ~ (e2 = op(e6, e0))) | (~ (e1 = op(e6, e6)) & ~ (e1 = op(e5, e6)) & ~ (e1 = op(e4, e6)) & ~ (e1 = op(e3, e6)) & ~ (e1 = op(e2, e6)) & ~ (e1 = op(e1, e6)) & ~ (e1 = op(e0, e6))) | (~ (e1 = op(e6, e6)) & ~ (e1 = op(e6, e5)) & ~ (e1 = op(e6, e4)) & ~ (e1 = op(e6, e3)) & ~ (e1 = op(e6, e2)) & ~ (e1 = op(e6, e1)) & ~ (e1 = op(e6, e0))) | (~ (e0 = op(e6, e6)) & ~ (e0 = op(e5, e6)) & ~ (e0 = op(e4, e6)) & ~ (e0 = op(e3, e6)) & ~ (e0 = op(e2, e6)) & ~ (e0 = op(e1, e6)) & ~ (e0 = op(e0, e6))) | (~ (e0 = op(e6, e6)) & ~ (e0 = op(e6, e5)) & ~ (e0 = op(e6, e4)) & ~ (e0 = op(e6, e3)) & ~ (e0 = op(e6, e2)) & ~ (e0 = op(e6, e1)) & ~ (e0 = op(e6, e0))) | (~ (e6 = op(e6, e5)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e0, e5))) | (~ (e6 = op(e5, e6)) & ~ (e6 = op(e5, e5)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e5, e0))) | (~ (e5 = op(e6, e5)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e0, e5))) | (~ (e5 = op(e5, e6)) & ~ (e5 = op(e5, e5)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e5, e0))) | (~ (e4 = op(e6, e5)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e0, e5))) | (~ (e4 = op(e5, e6)) & ~ (e4 = op(e5, e5)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e5, e0))) | (~ (e3 = op(e6, e5)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e0, e5))) | (~ (e3 = op(e5, e6)) & ~ (e3 = op(e5, e5)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e5, e0))) | (~ (e2 = op(e6, e5)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e0, e5))) | (~ (e2 = op(e5, e6)) & ~ (e2 = op(e5, e5)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e5, e0))) | (~ (e1 = op(e6, e5)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e0, e5))) | (~ (e1 = op(e5, e6)) & ~ (e1 = op(e5, e5)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e5, e0))) | (~ (e0 = op(e6, e5)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e0, e5))) | (~ (e0 = op(e5, e6)) & ~ (e0 = op(e5, e5)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e5, e0))) | (~ (e6 = op(e6, e4)) & ~ (e6 = op(e5, e4)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e0, e4))) | (~ (e6 = op(e4, e6)) & ~ (e6 = op(e4, e5)) & ~ (e6 = op(e4, e4)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e4, e0))) | (~ (e5 = op(e6, e4)) & ~ (e5 = op(e5, e4)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e0, e4))) | (~ (e5 = op(e4, e6)) & ~ (e5 = op(e4, e5)) & ~ (e5 = op(e4, e4)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e4, e0))) | (~ (e4 = op(e6, e4)) & ~ (e4 = op(e5, e4)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e0, e4))) | (~ (e4 = op(e4, e6)) & ~ (e4 = op(e4, e5)) & ~ (e4 = op(e4, e4)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e4, e0))) | (~ (e3 = op(e6, e4)) & ~ (e3 = op(e5, e4)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e0, e4))) | (~ (e3 = op(e4, e6)) & ~ (e3 = op(e4, e5)) & ~ (e3 = op(e4, e4)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e4, e0))) | (~ (e2 = op(e6, e4)) & ~ (e2 = op(e5, e4)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e0, e4))) | (~ (e2 = op(e4, e6)) & ~ (e2 = op(e4, e5)) & ~ (e2 = op(e4, e4)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e4, e0))) | (~ (e1 = op(e6, e4)) & ~ (e1 = op(e5, e4)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e0, e4))) | (~ (e1 = op(e4, e6)) & ~ (e1 = op(e4, e5)) & ~ (e1 = op(e4, e4)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e4, e0))) | (~ (e0 = op(e6, e4)) & ~ (e0 = op(e5, e4)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e0, e4))) | (~ (e0 = op(e4, e6)) & ~ (e0 = op(e4, e5)) & ~ (e0 = op(e4, e4)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e4, e0))) | (~ (e6 = op(e6, e3)) & ~ (e6 = op(e5, e3)) & ~ (e6 = op(e4, e3)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e0, e3))) | (~ (e6 = op(e3, e6)) & ~ (e6 = op(e3, e5)) & ~ (e6 = op(e3, e4)) & ~ (e6 = op(e3, e3)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e3, e0))) | (~ (e5 = op(e6, e3)) & ~ (e5 = op(e5, e3)) & ~ (e5 = op(e4, e3)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e0, e3))) | (~ (e5 = op(e3, e6)) & ~ (e5 = op(e3, e5)) & ~ (e5 = op(e3, e4)) & ~ (e5 = op(e3, e3)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e3, e0))) | (~ (e4 = op(e6, e3)) & ~ (e4 = op(e5, e3)) & ~ (e4 = op(e4, e3)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e0, e3))) | (~ (e4 = op(e3, e6)) & ~ (e4 = op(e3, e5)) & ~ (e4 = op(e3, e4)) & ~ (e4 = op(e3, e3)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e3, e0))) | (~ (e3 = op(e6, e3)) & ~ (e3 = op(e5, e3)) & ~ (e3 = op(e4, e3)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e0, e3))) | (~ (e3 = op(e3, e6)) & ~ (e3 = op(e3, e5)) & ~ (e3 = op(e3, e4)) & ~ (e3 = op(e3, e3)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e3, e0))) | (~ (e2 = op(e6, e3)) & ~ (e2 = op(e5, e3)) & ~ (e2 = op(e4, e3)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e0, e3))) | (~ (e2 = op(e3, e6)) & ~ (e2 = op(e3, e5)) & ~ (e2 = op(e3, e4)) & ~ (e2 = op(e3, e3)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e3, e0))) | (~ (e1 = op(e6, e3)) & ~ (e1 = op(e5, e3)) & ~ (e1 = op(e4, e3)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e0, e3))) | (~ (e1 = op(e3, e6)) & ~ (e1 = op(e3, e5)) & ~ (e1 = op(e3, e4)) & ~ (e1 = op(e3, e3)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e3, e0))) | (~ (e0 = op(e6, e3)) & ~ (e0 = op(e5, e3)) & ~ (e0 = op(e4, e3)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e0, e3))) | (~ (e0 = op(e3, e6)) & ~ (e0 = op(e3, e5)) & ~ (e0 = op(e3, e4)) & ~ (e0 = op(e3, e3)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e3, e0))) | (~ (e6 = op(e6, e2)) & ~ (e6 = op(e5, e2)) & ~ (e6 = op(e4, e2)) & ~ (e6 = op(e3, e2)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e0, e2))) | (~ (e6 = op(e2, e6)) & ~ (e6 = op(e2, e5)) & ~ (e6 = op(e2, e4)) & ~ (e6 = op(e2, e3)) & ~ (e6 = op(e2, e2)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e2, e0))) | (~ (e5 = op(e6, e2)) & ~ (e5 = op(e5, e2)) & ~ (e5 = op(e4, e2)) & ~ (e5 = op(e3, e2)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e0, e2))) | (~ (e5 = op(e2, e6)) & ~ (e5 = op(e2, e5)) & ~ (e5 = op(e2, e4)) & ~ (e5 = op(e2, e3)) & ~ (e5 = op(e2, e2)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e2, e0))) | (~ (e4 = op(e6, e2)) & ~ (e4 = op(e5, e2)) & ~ (e4 = op(e4, e2)) & ~ (e4 = op(e3, e2)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e0, e2))) | (~ (e4 = op(e2, e6)) & ~ (e4 = op(e2, e5)) & ~ (e4 = op(e2, e4)) & ~ (e4 = op(e2, e3)) & ~ (e4 = op(e2, e2)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e2, e0))) | (~ (e3 = op(e6, e2)) & ~ (e3 = op(e5, e2)) & ~ (e3 = op(e4, e2)) & ~ (e3 = op(e3, e2)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e0, e2))) | (~ (e3 = op(e2, e6)) & ~ (e3 = op(e2, e5)) & ~ (e3 = op(e2, e4)) & ~ (e3 = op(e2, e3)) & ~ (e3 = op(e2, e2)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e2, e0))) | (~ (e2 = op(e6, e2)) & ~ (e2 = op(e5, e2)) & ~ (e2 = op(e4, e2)) & ~ (e2 = op(e3, e2)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e0, e2))) | (~ (e2 = op(e2, e6)) & ~ (e2 = op(e2, e5)) & ~ (e2 = op(e2, e4)) & ~ (e2 = op(e2, e3)) & ~ (e2 = op(e2, e2)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e2, e0))) | (~ (e1 = op(e6, e2)) & ~ (e1 = op(e5, e2)) & ~ (e1 = op(e4, e2)) & ~ (e1 = op(e3, e2)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e0, e2))) | (~ (e1 = op(e2, e6)) & ~ (e1 = op(e2, e5)) & ~ (e1 = op(e2, e4)) & ~ (e1 = op(e2, e3)) & ~ (e1 = op(e2, e2)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e2, e0))) | (~ (e0 = op(e6, e2)) & ~ (e0 = op(e5, e2)) & ~ (e0 = op(e4, e2)) & ~ (e0 = op(e3, e2)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e0, e2))) | (~ (e0 = op(e2, e6)) & ~ (e0 = op(e2, e5)) & ~ (e0 = op(e2, e4)) & ~ (e0 = op(e2, e3)) & ~ (e0 = op(e2, e2)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e2, e0))) | (~ (e6 = op(e6, e1)) & ~ (e6 = op(e5, e1)) & ~ (e6 = op(e4, e1)) & ~ (e6 = op(e3, e1)) & ~ (e6 = op(e2, e1)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e0, e1))) | (~ (e6 = op(e1, e6)) & ~ (e6 = op(e1, e5)) & ~ (e6 = op(e1, e4)) & ~ (e6 = op(e1, e3)) & ~ (e6 = op(e1, e2)) & ~ (e6 = op(e1, e1)) & ~ (e6 = op(e1, e0))) | (~ (e5 = op(e6, e1)) & ~ (e5 = op(e5, e1)) & ~ (e5 = op(e4, e1)) & ~ (e5 = op(e3, e1)) & ~ (e5 = op(e2, e1)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e0, e1))) | (~ (e5 = op(e1, e6)) & ~ (e5 = op(e1, e5)) & ~ (e5 = op(e1, e4)) & ~ (e5 = op(e1, e3)) & ~ (e5 = op(e1, e2)) & ~ (e5 = op(e1, e1)) & ~ (e5 = op(e1, e0))) | (~ (e4 = op(e6, e1)) & ~ (e4 = op(e5, e1)) & ~ (e4 = op(e4, e1)) & ~ (e4 = op(e3, e1)) & ~ (e4 = op(e2, e1)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e0, e1))) | (~ (e4 = op(e1, e6)) & ~ (e4 = op(e1, e5)) & ~ (e4 = op(e1, e4)) & ~ (e4 = op(e1, e3)) & ~ (e4 = op(e1, e2)) & ~ (e4 = op(e1, e1)) & ~ (e4 = op(e1, e0))) | (~ (e3 = op(e6, e1)) & ~ (e3 = op(e5, e1)) & ~ (e3 = op(e4, e1)) & ~ (e3 = op(e3, e1)) & ~ (e3 = op(e2, e1)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e0, e1))) | (~ (e3 = op(e1, e6)) & ~ (e3 = op(e1, e5)) & ~ (e3 = op(e1, e4)) & ~ (e3 = op(e1, e3)) & ~ (e3 = op(e1, e2)) & ~ (e3 = op(e1, e1)) & ~ (e3 = op(e1, e0))) | (~ (e2 = op(e6, e1)) & ~ (e2 = op(e5, e1)) & ~ (e2 = op(e4, e1)) & ~ (e2 = op(e3, e1)) & ~ (e2 = op(e2, e1)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e0, e1))) | (~ (e2 = op(e1, e6)) & ~ (e2 = op(e1, e5)) & ~ (e2 = op(e1, e4)) & ~ (e2 = op(e1, e3)) & ~ (e2 = op(e1, e2)) & ~ (e2 = op(e1, e1)) & ~ (e2 = op(e1, e0))) | (~ (e1 = op(e6, e1)) & ~ (e1 = op(e5, e1)) & ~ (e1 = op(e4, e1)) & ~ (e1 = op(e3, e1)) & ~ (e1 = op(e2, e1)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e0, e1))) | (~ (e1 = op(e1, e6)) & ~ (e1 = op(e1, e5)) & ~ (e1 = op(e1, e4)) & ~ (e1 = op(e1, e3)) & ~ (e1 = op(e1, e2)) & ~ (e1 = op(e1, e1)) & ~ (e1 = op(e1, e0))) | (~ (e0 = op(e6, e1)) & ~ (e0 = op(e5, e1)) & ~ (e0 = op(e4, e1)) & ~ (e0 = op(e3, e1)) & ~ (e0 = op(e2, e1)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e0, e1))) | (~ (e0 = op(e1, e6)) & ~ (e0 = op(e1, e5)) & ~ (e0 = op(e1, e4)) & ~ (e0 = op(e1, e3)) & ~ (e0 = op(e1, e2)) & ~ (e0 = op(e1, e1)) & ~ (e0 = op(e1, e0))) | (~ (e6 = op(e6, e0)) & ~ (e6 = op(e5, e0)) & ~ (e6 = op(e4, e0)) & ~ (e6 = op(e3, e0)) & ~ (e6 = op(e2, e0)) & ~ (e6 = op(e1, e0)) & ~ (e6 = op(e0, e0))) | (~ (e6 = op(e0, e6)) & ~ (e6 = op(e0, e5)) & ~ (e6 = op(e0, e4)) & ~ (e6 = op(e0, e3)) & ~ (e6 = op(e0, e2)) & ~ (e6 = op(e0, e1)) & ~ (e6 = op(e0, e0))) | (~ (e5 = op(e6, e0)) & ~ (e5 = op(e5, e0)) & ~ (e5 = op(e4, e0)) & ~ (e5 = op(e3, e0)) & ~ (e5 = op(e2, e0)) & ~ (e5 = op(e1, e0)) & ~ (e5 = op(e0, e0))) | (~ (e5 = op(e0, e6)) & ~ (e5 = op(e0, e5)) & ~ (e5 = op(e0, e4)) & ~ (e5 = op(e0, e3)) & ~ (e5 = op(e0, e2)) & ~ (e5 = op(e0, e1)) & ~ (e5 = op(e0, e0))) | (~ (e4 = op(e6, e0)) & ~ (e4 = op(e5, e0)) & ~ (e4 = op(e4, e0)) & ~ (e4 = op(e3, e0)) & ~ (e4 = op(e2, e0)) & ~ (e4 = op(e1, e0)) & ~ (e4 = op(e0, e0))) | (~ (e4 = op(e0, e6)) & ~ (e4 = op(e0, e5)) & ~ (e4 = op(e0, e4)) & ~ (e4 = op(e0, e3)) & ~ (e4 = op(e0, e2)) & ~ (e4 = op(e0, e1)) & ~ (e4 = op(e0, e0))) | (~ (e3 = op(e6, e0)) & ~ (e3 = op(e5, e0)) & ~ (e3 = op(e4, e0)) & ~ (e3 = op(e3, e0)) & ~ (e3 = op(e2, e0)) & ~ (e3 = op(e1, e0)) & ~ (e3 = op(e0, e0))) | (~ (e3 = op(e0, e6)) & ~ (e3 = op(e0, e5)) & ~ (e3 = op(e0, e4)) & ~ (e3 = op(e0, e3)) & ~ (e3 = op(e0, e2)) & ~ (e3 = op(e0, e1)) & ~ (e3 = op(e0, e0))) | (~ (e2 = op(e6, e0)) & ~ (e2 = op(e5, e0)) & ~ (e2 = op(e4, e0)) & ~ (e2 = op(e3, e0)) & ~ (e2 = op(e2, e0)) & ~ (e2 = op(e1, e0)) & ~ (e2 = op(e0, e0))) | (~ (e2 = op(e0, e6)) & ~ (e2 = op(e0, e5)) & ~ (e2 = op(e0, e4)) & ~ (e2 = op(e0, e3)) & ~ (e2 = op(e0, e2)) & ~ (e2 = op(e0, e1)) & ~ (e2 = op(e0, e0))) | (~ (e1 = op(e6, e0)) & ~ (e1 = op(e5, e0)) & ~ (e1 = op(e4, e0)) & ~ (e1 = op(e3, e0)) & ~ (e1 = op(e2, e0)) & ~ (e1 = op(e1, e0)) & ~ (e1 = op(e0, e0))) | (~ (e1 = op(e0, e6)) & ~ (e1 = op(e0, e5)) & ~ (e1 = op(e0, e4)) & ~ (e1 = op(e0, e3)) & ~ (e1 = op(e0, e2)) & ~ (e1 = op(e0, e1)) & ~ (e1 = op(e0, e0))) | (~ (e0 = op(e6, e0)) & ~ (e0 = op(e5, e0)) & ~ (e0 = op(e4, e0)) & ~ (e0 = op(e3, e0)) & ~ (e0 = op(e2, e0)) & ~ (e0 = op(e1, e0)) & ~ (e0 = op(e0, e0))) | (~ (e0 = op(e0, e6)) & ~ (e0 = op(e0, e5)) & ~ (e0 = op(e0, e4)) & ~ (e0 = op(e0, e3)) & ~ (e0 = op(e0, e2)) & ~ (e0 = op(e0, e1)) & ~ (e0 = op(e0, e0))) | (~ (e6 = op(e6, e6)) & ~ (e5 = op(e6, e6)) & ~ (e4 = op(e6, e6)) & ~ (e3 = op(e6, e6)) & ~ (e2 = op(e6, e6)) & ~ (e1 = op(e6, e6)) & ~ (e0 = op(e6, e6))) | (~ (e6 = op(e6, e5)) & ~ (e5 = op(e6, e5)) & ~ (e4 = op(e6, e5)) & ~ (e3 = op(e6, e5)) & ~ (e2 = op(e6, e5)) & ~ (e1 = op(e6, e5)) & ~ (e0 = op(e6, e5))) | (~ (e6 = op(e6, e4)) & ~ (e5 = op(e6, e4)) & ~ (e4 = op(e6, e4)) & ~ (e3 = op(e6, e4)) & ~ (e2 = op(e6, e4)) & ~ (e1 = op(e6, e4)) & ~ (e0 = op(e6, e4))) | (~ (e6 = op(e6, e3)) & ~ (e5 = op(e6, e3)) & ~ (e4 = op(e6, e3)) & ~ (e3 = op(e6, e3)) & ~ (e2 = op(e6, e3)) & ~ (e1 = op(e6, e3)) & ~ (e0 = op(e6, e3))) | (~ (e6 = op(e6, e2)) & ~ (e5 = op(e6, e2)) & ~ (e4 = op(e6, e2)) & ~ (e3 = op(e6, e2)) & ~ (e2 = op(e6, e2)) & ~ (e1 = op(e6, e2)) & ~ (e0 = op(e6, e2))) | (~ (e6 = op(e6, e1)) & ~ (e5 = op(e6, e1)) & ~ (e4 = op(e6, e1)) & ~ (e3 = op(e6, e1)) & ~ (e2 = op(e6, e1)) & ~ (e1 = op(e6, e1)) & ~ (e0 = op(e6, e1))) | (~ (e6 = op(e6, e0)) & ~ (e5 = op(e6, e0)) & ~ (e4 = op(e6, e0)) & ~ (e3 = op(e6, e0)) & ~ (e2 = op(e6, e0)) & ~ (e1 = op(e6, e0)) & ~ (e0 = op(e6, e0))) | (~ (e6 = op(e5, e6)) & ~ (e5 = op(e5, e6)) & ~ (e4 = op(e5, e6)) & ~ (e3 = op(e5, e6)) & ~ (e2 = op(e5, e6)) & ~ (e1 = op(e5, e6)) & ~ (e0 = op(e5, e6))) | (~ (e6 = op(e5, e5)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e5, e5)) & ~ (e3 = op(e5, e5)) & ~ (e2 = op(e5, e5)) & ~ (e1 = op(e5, e5)) & ~ (e0 = op(e5, e5))) | (~ (e6 = op(e5, e4)) & ~ (e5 = op(e5, e4)) & ~ (e4 = op(e5, e4)) & ~ (e3 = op(e5, e4)) & ~ (e2 = op(e5, e4)) & ~ (e1 = op(e5, e4)) & ~ (e0 = op(e5, e4))) | (~ (e6 = op(e5, e3)) & ~ (e5 = op(e5, e3)) & ~ (e4 = op(e5, e3)) & ~ (e3 = op(e5, e3)) & ~ (e2 = op(e5, e3)) & ~ (e1 = op(e5, e3)) & ~ (e0 = op(e5, e3))) | (~ (e6 = op(e5, e2)) & ~ (e5 = op(e5, e2)) & ~ (e4 = op(e5, e2)) & ~ (e3 = op(e5, e2)) & ~ (e2 = op(e5, e2)) & ~ (e1 = op(e5, e2)) & ~ (e0 = op(e5, e2))) | (~ (e6 = op(e5, e1)) & ~ (e5 = op(e5, e1)) & ~ (e4 = op(e5, e1)) & ~ (e3 = op(e5, e1)) & ~ (e2 = op(e5, e1)) & ~ (e1 = op(e5, e1)) & ~ (e0 = op(e5, e1))) | (~ (e6 = op(e5, e0)) & ~ (e5 = op(e5, e0)) & ~ (e4 = op(e5, e0)) & ~ (e3 = op(e5, e0)) & ~ (e2 = op(e5, e0)) & ~ (e1 = op(e5, e0)) & ~ (e0 = op(e5, e0))) | (~ (e6 = op(e4, e6)) & ~ (e5 = op(e4, e6)) & ~ (e4 = op(e4, e6)) & ~ (e3 = op(e4, e6)) & ~ (e2 = op(e4, e6)) & ~ (e1 = op(e4, e6)) & ~ (e0 = op(e4, e6))) | (~ (e6 = op(e4, e5)) & ~ (e5 = op(e4, e5)) & ~ (e4 = op(e4, e5)) & ~ (e3 = op(e4, e5)) & ~ (e2 = op(e4, e5)) & ~ (e1 = op(e4, e5)) & ~ (e0 = op(e4, e5))) | (~ (e6 = op(e4, e4)) & ~ (e5 = op(e4, e4)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e4, e4)) & ~ (e2 = op(e4, e4)) & ~ (e1 = op(e4, e4)) & ~ (e0 = op(e4, e4))) | (~ (e6 = op(e4, e3)) & ~ (e5 = op(e4, e3)) & ~ (e4 = op(e4, e3)) & ~ (e3 = op(e4, e3)) & ~ (e2 = op(e4, e3)) & ~ (e1 = op(e4, e3)) & ~ (e0 = op(e4, e3))) | (~ (e6 = op(e4, e2)) & ~ (e5 = op(e4, e2)) & ~ (e4 = op(e4, e2)) & ~ (e3 = op(e4, e2)) & ~ (e2 = op(e4, e2)) & ~ (e1 = op(e4, e2)) & ~ (e0 = op(e4, e2))) | (~ (e6 = op(e4, e1)) & ~ (e5 = op(e4, e1)) & ~ (e4 = op(e4, e1)) & ~ (e3 = op(e4, e1)) & ~ (e2 = op(e4, e1)) & ~ (e1 = op(e4, e1)) & ~ (e0 = op(e4, e1))) | (~ (e6 = op(e4, e0)) & ~ (e5 = op(e4, e0)) & ~ (e4 = op(e4, e0)) & ~ (e3 = op(e4, e0)) & ~ (e2 = op(e4, e0)) & ~ (e1 = op(e4, e0)) & ~ (e0 = op(e4, e0))) | (~ (e6 = op(e3, e6)) & ~ (e5 = op(e3, e6)) & ~ (e4 = op(e3, e6)) & ~ (e3 = op(e3, e6)) & ~ (e2 = op(e3, e6)) & ~ (e1 = op(e3, e6)) & ~ (e0 = op(e3, e6))) | (~ (e6 = op(e3, e5)) & ~ (e5 = op(e3, e5)) & ~ (e4 = op(e3, e5)) & ~ (e3 = op(e3, e5)) & ~ (e2 = op(e3, e5)) & ~ (e1 = op(e3, e5)) & ~ (e0 = op(e3, e5))) | (~ (e6 = op(e3, e4)) & ~ (e5 = op(e3, e4)) & ~ (e4 = op(e3, e4)) & ~ (e3 = op(e3, e4)) & ~ (e2 = op(e3, e4)) & ~ (e1 = op(e3, e4)) & ~ (e0 = op(e3, e4))) | (~ (e6 = op(e3, e3)) & ~ (e5 = op(e3, e3)) & ~ (e4 = op(e3, e3)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e3, e3)) & ~ (e1 = op(e3, e3)) & ~ (e0 = op(e3, e3))) | (~ (e6 = op(e3, e2)) & ~ (e5 = op(e3, e2)) & ~ (e4 = op(e3, e2)) & ~ (e3 = op(e3, e2)) & ~ (e2 = op(e3, e2)) & ~ (e1 = op(e3, e2)) & ~ (e0 = op(e3, e2))) | (~ (e6 = op(e3, e1)) & ~ (e5 = op(e3, e1)) & ~ (e4 = op(e3, e1)) & ~ (e3 = op(e3, e1)) & ~ (e2 = op(e3, e1)) & ~ (e1 = op(e3, e1)) & ~ (e0 = op(e3, e1))) | (~ (e6 = op(e3, e0)) & ~ (e5 = op(e3, e0)) & ~ (e4 = op(e3, e0)) & ~ (e3 = op(e3, e0)) & ~ (e2 = op(e3, e0)) & ~ (e1 = op(e3, e0)) & ~ (e0 = op(e3, e0))) | (~ (e6 = op(e2, e6)) & ~ (e5 = op(e2, e6)) & ~ (e4 = op(e2, e6)) & ~ (e3 = op(e2, e6)) & ~ (e2 = op(e2, e6)) & ~ (e1 = op(e2, e6)) & ~ (e0 = op(e2, e6))) | (~ (e6 = op(e2, e5)) & ~ (e5 = op(e2, e5)) & ~ (e4 = op(e2, e5)) & ~ (e3 = op(e2, e5)) & ~ (e2 = op(e2, e5)) & ~ (e1 = op(e2, e5)) & ~ (e0 = op(e2, e5))) | (~ (e6 = op(e2, e4)) & ~ (e5 = op(e2, e4)) & ~ (e4 = op(e2, e4)) & ~ (e3 = op(e2, e4)) & ~ (e2 = op(e2, e4)) & ~ (e1 = op(e2, e4)) & ~ (e0 = op(e2, e4))) | (~ (e6 = op(e2, e3)) & ~ (e5 = op(e2, e3)) & ~ (e4 = op(e2, e3)) & ~ (e3 = op(e2, e3)) & ~ (e2 = op(e2, e3)) & ~ (e1 = op(e2, e3)) & ~ (e0 = op(e2, e3))) | (~ (e6 = op(e2, e2)) & ~ (e5 = op(e2, e2)) & ~ (e4 = op(e2, e2)) & ~ (e3 = op(e2, e2)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e2, e2)) & ~ (e0 = op(e2, e2))) | (~ (e6 = op(e2, e1)) & ~ (e5 = op(e2, e1)) & ~ (e4 = op(e2, e1)) & ~ (e3 = op(e2, e1)) & ~ (e2 = op(e2, e1)) & ~ (e1 = op(e2, e1)) & ~ (e0 = op(e2, e1))) | (~ (e6 = op(e2, e0)) & ~ (e5 = op(e2, e0)) & ~ (e4 = op(e2, e0)) & ~ (e3 = op(e2, e0)) & ~ (e2 = op(e2, e0)) & ~ (e1 = op(e2, e0)) & ~ (e0 = op(e2, e0))) | (~ (e6 = op(e1, e6)) & ~ (e5 = op(e1, e6)) & ~ (e4 = op(e1, e6)) & ~ (e3 = op(e1, e6)) & ~ (e2 = op(e1, e6)) & ~ (e1 = op(e1, e6)) & ~ (e0 = op(e1, e6))) | (~ (e6 = op(e1, e5)) & ~ (e5 = op(e1, e5)) & ~ (e4 = op(e1, e5)) & ~ (e3 = op(e1, e5)) & ~ (e2 = op(e1, e5)) & ~ (e1 = op(e1, e5)) & ~ (e0 = op(e1, e5))) | (~ (e6 = op(e1, e4)) & ~ (e5 = op(e1, e4)) & ~ (e4 = op(e1, e4)) & ~ (e3 = op(e1, e4)) & ~ (e2 = op(e1, e4)) & ~ (e1 = op(e1, e4)) & ~ (e0 = op(e1, e4))) | (~ (e6 = op(e1, e3)) & ~ (e5 = op(e1, e3)) & ~ (e4 = op(e1, e3)) & ~ (e3 = op(e1, e3)) & ~ (e2 = op(e1, e3)) & ~ (e1 = op(e1, e3)) & ~ (e0 = op(e1, e3))) | (~ (e6 = op(e1, e2)) & ~ (e5 = op(e1, e2)) & ~ (e4 = op(e1, e2)) & ~ (e3 = op(e1, e2)) & ~ (e2 = op(e1, e2)) & ~ (e1 = op(e1, e2)) & ~ (e0 = op(e1, e2))) | (~ (e6 = op(e1, e1)) & ~ (e5 = op(e1, e1)) & ~ (e4 = op(e1, e1)) & ~ (e3 = op(e1, e1)) & ~ (e2 = op(e1, e1)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e1, e1))) | (~ (e6 = op(e1, e0)) & ~ (e5 = op(e1, e0)) & ~ (e4 = op(e1, e0)) & ~ (e3 = op(e1, e0)) & ~ (e2 = op(e1, e0)) & ~ (e1 = op(e1, e0)) & ~ (e0 = op(e1, e0))) | (~ (e6 = op(e0, e6)) & ~ (e5 = op(e0, e6)) & ~ (e4 = op(e0, e6)) & ~ (e3 = op(e0, e6)) & ~ (e2 = op(e0, e6)) & ~ (e1 = op(e0, e6)) & ~ (e0 = op(e0, e6))) | (~ (e6 = op(e0, e5)) & ~ (e5 = op(e0, e5)) & ~ (e4 = op(e0, e5)) & ~ (e3 = op(e0, e5)) & ~ (e2 = op(e0, e5)) & ~ (e1 = op(e0, e5)) & ~ (e0 = op(e0, e5))) | (~ (e6 = op(e0, e4)) & ~ (e5 = op(e0, e4)) & ~ (e4 = op(e0, e4)) & ~ (e3 = op(e0, e4)) & ~ (e2 = op(e0, e4)) & ~ (e1 = op(e0, e4)) & ~ (e0 = op(e0, e4))) | (~ (e6 = op(e0, e3)) & ~ (e5 = op(e0, e3)) & ~ (e4 = op(e0, e3)) & ~ (e3 = op(e0, e3)) & ~ (e2 = op(e0, e3)) & ~ (e1 = op(e0, e3)) & ~ (e0 = op(e0, e3))) | (~ (e6 = op(e0, e2)) & ~ (e5 = op(e0, e2)) & ~ (e4 = op(e0, e2)) & ~ (e3 = op(e0, e2)) & ~ (e2 = op(e0, e2)) & ~ (e1 = op(e0, e2)) & ~ (e0 = op(e0, e2))) | (~ (e6 = op(e0, e1)) & ~ (e5 = op(e0, e1)) & ~ (e4 = op(e0, e1)) & ~ (e3 = op(e0, e1)) & ~ (e2 = op(e0, e1)) & ~ (e1 = op(e0, e1)) & ~ (e0 = op(e0, e1))) | (~ (e6 = op(e0, e0)) & ~ (e5 = op(e0, e0)) & ~ (e4 = op(e0, e0)) & ~ (e3 = op(e0, e0)) & ~ (e2 = op(e0, e0)) & ~ (e1 = op(e0, e0)) & ~ (e0 = op(e0, e0))) | ((e6 = op(e6, e6)) & (e5 = op(e5, e5)) & (e4 = op(e4, e4)) & (e3 = op(e3, e3)) & (e2 = op(e2, e2)) & (e1 = op(e1, e1)) & (e0 = op(e0, e0))) | (~ (e6 = op(e6, e6)) & ~ (e5 = op(e5, e5)) & ~ (e4 = op(e4, e4)) & ~ (e3 = op(e3, e3)) & ~ (e2 = op(e2, e2)) & ~ (e1 = op(e1, e1)) & ~ (e0 = op(e0, e0)))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ~ ((e6 = op(op(op(e6, e6), e6), e6)) & (e6 = op(op(op(e5, e6), e5), e5)) & (e6 = op(op(op(e4, e6), e4), e4)) & (e6 = op(op(op(e3, e6), e3), e3)) & (e6 = op(op(op(e2, e6), e2), e2)) & (e6 = op(op(op(e1, e6), e1), e1)) & (e6 = op(op(op(e0, e6), e0), e0)) & (e5 = op(op(op(e6, e5), e6), e6)) & (e5 = op(op(op(e5, e5), e5), e5)) & (e5 = op(op(op(e4, e5), e4), e4)) & (e5 = op(op(op(e3, e5), e3), e3)) & (e5 = op(op(op(e2, e5), e2), e2)) & (e5 = op(op(op(e1, e5), e1), e1)) & (e5 = op(op(op(e0, e5), e0), e0)) & (e4 = op(op(op(e6, e4), e6), e6)) & (e4 = op(op(op(e5, e4), e5), e5)) & (e4 = op(op(op(e4, e4), e4), e4)) & (e4 = op(op(op(e3, e4), e3), e3)) & (e4 = op(op(op(e2, e4), e2), e2)) & (e4 = op(op(op(e1, e4), e1), e1)) & (e4 = op(op(op(e0, e4), e0), e0)) & (e3 = op(op(op(e6, e3), e6), e6)) & (e3 = op(op(op(e5, e3), e5), e5)) & (e3 = op(op(op(e4, e3), e4), e4)) & (e3 = op(op(op(e3, e3), e3), e3)) & (e3 = op(op(op(e2, e3), e2), e2)) & (e3 = op(op(op(e1, e3), e1), e1)) & (e3 = op(op(op(e0, e3), e0), e0)) & (e2 = op(op(op(e6, e2), e6), e6)) & (e2 = op(op(op(e5, e2), e5), e5)) & (e2 = op(op(op(e4, e2), e4), e4)) & (e2 = op(op(op(e3, e2), e3), e3)) & (e2 = op(op(op(e2, e2), e2), e2)) & (e2 = op(op(op(e1, e2), e1), e1)) & (e2 = op(op(op(e0, e2), e0), e0)) & (e1 = op(op(op(e6, e1), e6), e6)) & (e1 = op(op(op(e5, e1), e5), e5)) & (e1 = op(op(op(e4, e1), e4), e4)) & (e1 = op(op(op(e3, e1), e3), e3)) & (e1 = op(op(op(e2, e1), e2), e2)) & (e1 = op(op(op(e1, e1), e1), e1)) & (e1 = op(op(op(e0, e1), e0), e0)) & (e0 = op(op(op(e6, e0), e6), e6)) & (e0 = op(op(op(e5, e0), e5), e5)) & (e0 = op(op(op(e4, e0), e4), e4)) & (e0 = op(op(op(e3, e0), e3), e3)) & (e0 = op(op(op(e2, e0), e2), e2)) & (e0 = op(op(op(e1, e0), e1), e1)) & (e0 = op(op(op(e0, e0), e0), e0)) & ((e6 = op(e6, e6)) | (e6 = op(e5, e6)) | (e6 = op(e4, e6)) | (e6 = op(e3, e6)) | (e6 = op(e2, e6)) | (e6 = op(e1, e6)) | (e6 = op(e0, e6))) & ((e6 = op(e6, e6)) | (e6 = op(e6, e5)) | (e6 = op(e6, e4)) | (e6 = op(e6, e3)) | (e6 = op(e6, e2)) | (e6 = op(e6, e1)) | (e6 = op(e6, e0))) & ((e5 = op(e6, e6)) | (e5 = op(e5, e6)) | (e5 = op(e4, e6)) | (e5 = op(e3, e6)) | (e5 = op(e2, e6)) | (e5 = op(e1, e6)) | (e5 = op(e0, e6))) & ((e5 = op(e6, e6)) | (e5 = op(e6, e5)) | (e5 = op(e6, e4)) | (e5 = op(e6, e3)) | (e5 = op(e6, e2)) | (e5 = op(e6, e1)) | (e5 = op(e6, e0))) & ((e4 = op(e6, e6)) | (e4 = op(e5, e6)) | (e4 = op(e4, e6)) | (e4 = op(e3, e6)) | (e4 = op(e2, e6)) | (e4 = op(e1, e6)) | (e4 = op(e0, e6))) & ((e4 = op(e6, e6)) | (e4 = op(e6, e5)) | (e4 = op(e6, e4)) | (e4 = op(e6, e3)) | (e4 = op(e6, e2)) | (e4 = op(e6, e1)) | (e4 = op(e6, e0))) & ((e3 = op(e6, e6)) | (e3 = op(e5, e6)) | (e3 = op(e4, e6)) | (e3 = op(e3, e6)) | (e3 = op(e2, e6)) | (e3 = op(e1, e6)) | (e3 = op(e0, e6))) & ((e3 = op(e6, e6)) | (e3 = op(e6, e5)) | (e3 = op(e6, e4)) | (e3 = op(e6, e3)) | (e3 = op(e6, e2)) | (e3 = op(e6, e1)) | (e3 = op(e6, e0))) & ((e2 = op(e6, e6)) | (e2 = op(e5, e6)) | (e2 = op(e4, e6)) | (e2 = op(e3, e6)) | (e2 = op(e2, e6)) | (e2 = op(e1, e6)) | (e2 = op(e0, e6))) & ((e2 = op(e6, e6)) | (e2 = op(e6, e5)) | (e2 = op(e6, e4)) | (e2 = op(e6, e3)) | (e2 = op(e6, e2)) | (e2 = op(e6, e1)) | (e2 = op(e6, e0))) & ((e1 = op(e6, e6)) | (e1 = op(e5, e6)) | (e1 = op(e4, e6)) | (e1 = op(e3, e6)) | (e1 = op(e2, e6)) | (e1 = op(e1, e6)) | (e1 = op(e0, e6))) & ((e1 = op(e6, e6)) | (e1 = op(e6, e5)) | (e1 = op(e6, e4)) | (e1 = op(e6, e3)) | (e1 = op(e6, e2)) | (e1 = op(e6, e1)) | (e1 = op(e6, e0))) & ((e0 = op(e6, e6)) | (e0 = op(e5, e6)) | (e0 = op(e4, e6)) | (e0 = op(e3, e6)) | (e0 = op(e2, e6)) | (e0 = op(e1, e6)) | (e0 = op(e0, e6))) & ((e0 = op(e6, e6)) | (e0 = op(e6, e5)) | (e0 = op(e6, e4)) | (e0 = op(e6, e3)) | (e0 = op(e6, e2)) | (e0 = op(e6, e1)) | (e0 = op(e6, e0))) & ((e6 = op(e6, e5)) | (e6 = op(e5, e5)) | (e6 = op(e4, e5)) | (e6 = op(e3, e5)) | (e6 = op(e2, e5)) | (e6 = op(e1, e5)) | (e6 = op(e0, e5))) & ((e6 = op(e5, e6)) | (e6 = op(e5, e5)) | (e6 = op(e5, e4)) | (e6 = op(e5, e3)) | (e6 = op(e5, e2)) | (e6 = op(e5, e1)) | (e6 = op(e5, e0))) & ((e5 = op(e6, e5)) | (e5 = op(e5, e5)) | (e5 = op(e4, e5)) | (e5 = op(e3, e5)) | (e5 = op(e2, e5)) | (e5 = op(e1, e5)) | (e5 = op(e0, e5))) & ((e5 = op(e5, e6)) | (e5 = op(e5, e5)) | (e5 = op(e5, e4)) | (e5 = op(e5, e3)) | (e5 = op(e5, e2)) | (e5 = op(e5, e1)) | (e5 = op(e5, e0))) & ((e4 = op(e6, e5)) | (e4 = op(e5, e5)) | (e4 = op(e4, e5)) | (e4 = op(e3, e5)) | (e4 = op(e2, e5)) | (e4 = op(e1, e5)) | (e4 = op(e0, e5))) & ((e4 = op(e5, e6)) | (e4 = op(e5, e5)) | (e4 = op(e5, e4)) | (e4 = op(e5, e3)) | (e4 = op(e5, e2)) | (e4 = op(e5, e1)) | (e4 = op(e5, e0))) & ((e3 = op(e6, e5)) | (e3 = op(e5, e5)) | (e3 = op(e4, e5)) | (e3 = op(e3, e5)) | (e3 = op(e2, e5)) | (e3 = op(e1, e5)) | (e3 = op(e0, e5))) & ((e3 = op(e5, e6)) | (e3 = op(e5, e5)) | (e3 = op(e5, e4)) | (e3 = op(e5, e3)) | (e3 = op(e5, e2)) | (e3 = op(e5, e1)) | (e3 = op(e5, e0))) & ((e2 = op(e6, e5)) | (e2 = op(e5, e5)) | (e2 = op(e4, e5)) | (e2 = op(e3, e5)) | (e2 = op(e2, e5)) | (e2 = op(e1, e5)) | (e2 = op(e0, e5))) & ((e2 = op(e5, e6)) | (e2 = op(e5, e5)) | (e2 = op(e5, e4)) | (e2 = op(e5, e3)) | (e2 = op(e5, e2)) | (e2 = op(e5, e1)) | (e2 = op(e5, e0))) & ((e1 = op(e6, e5)) | (e1 = op(e5, e5)) | (e1 = op(e4, e5)) | (e1 = op(e3, e5)) | (e1 = op(e2, e5)) | (e1 = op(e1, e5)) | (e1 = op(e0, e5))) & ((e1 = op(e5, e6)) | (e1 = op(e5, e5)) | (e1 = op(e5, e4)) | (e1 = op(e5, e3)) | (e1 = op(e5, e2)) | (e1 = op(e5, e1)) | (e1 = op(e5, e0))) & ((e0 = op(e6, e5)) | (e0 = op(e5, e5)) | (e0 = op(e4, e5)) | (e0 = op(e3, e5)) | (e0 = op(e2, e5)) | (e0 = op(e1, e5)) | (e0 = op(e0, e5))) & ((e0 = op(e5, e6)) | (e0 = op(e5, e5)) | (e0 = op(e5, e4)) | (e0 = op(e5, e3)) | (e0 = op(e5, e2)) | (e0 = op(e5, e1)) | (e0 = op(e5, e0))) & ((e6 = op(e6, e4)) | (e6 = op(e5, e4)) | (e6 = op(e4, e4)) | (e6 = op(e3, e4)) | (e6 = op(e2, e4)) | (e6 = op(e1, e4)) | (e6 = op(e0, e4))) & ((e6 = op(e4, e6)) | (e6 = op(e4, e5)) | (e6 = op(e4, e4)) | (e6 = op(e4, e3)) | (e6 = op(e4, e2)) | (e6 = op(e4, e1)) | (e6 = op(e4, e0))) & ((e5 = op(e6, e4)) | (e5 = op(e5, e4)) | (e5 = op(e4, e4)) | (e5 = op(e3, e4)) | (e5 = op(e2, e4)) | (e5 = op(e1, e4)) | (e5 = op(e0, e4))) & ((e5 = op(e4, e6)) | (e5 = op(e4, e5)) | (e5 = op(e4, e4)) | (e5 = op(e4, e3)) | (e5 = op(e4, e2)) | (e5 = op(e4, e1)) | (e5 = op(e4, e0))) & ((e4 = op(e6, e4)) | (e4 = op(e5, e4)) | (e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e6)) | (e4 = op(e4, e5)) | (e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e6, e4)) | (e3 = op(e5, e4)) | (e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e6)) | (e3 = op(e4, e5)) | (e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e6, e4)) | (e2 = op(e5, e4)) | (e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e6)) | (e2 = op(e4, e5)) | (e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e6, e4)) | (e1 = op(e5, e4)) | (e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e6)) | (e1 = op(e4, e5)) | (e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e6, e4)) | (e0 = op(e5, e4)) | (e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e6)) | (e0 = op(e4, e5)) | (e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e6 = op(e6, e3)) | (e6 = op(e5, e3)) | (e6 = op(e4, e3)) | (e6 = op(e3, e3)) | (e6 = op(e2, e3)) | (e6 = op(e1, e3)) | (e6 = op(e0, e3))) & ((e6 = op(e3, e6)) | (e6 = op(e3, e5)) | (e6 = op(e3, e4)) | (e6 = op(e3, e3)) | (e6 = op(e3, e2)) | (e6 = op(e3, e1)) | (e6 = op(e3, e0))) & ((e5 = op(e6, e3)) | (e5 = op(e5, e3)) | (e5 = op(e4, e3)) | (e5 = op(e3, e3)) | (e5 = op(e2, e3)) | (e5 = op(e1, e3)) | (e5 = op(e0, e3))) & ((e5 = op(e3, e6)) | (e5 = op(e3, e5)) | (e5 = op(e3, e4)) | (e5 = op(e3, e3)) | (e5 = op(e3, e2)) | (e5 = op(e3, e1)) | (e5 = op(e3, e0))) & ((e4 = op(e6, e3)) | (e4 = op(e5, e3)) | (e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e6)) | (e4 = op(e3, e5)) | (e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e6, e3)) | (e3 = op(e5, e3)) | (e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e6)) | (e3 = op(e3, e5)) | (e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e6, e3)) | (e2 = op(e5, e3)) | (e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e6)) | (e2 = op(e3, e5)) | (e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e6, e3)) | (e1 = op(e5, e3)) | (e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e6)) | (e1 = op(e3, e5)) | (e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e6, e3)) | (e0 = op(e5, e3)) | (e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e6)) | (e0 = op(e3, e5)) | (e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e6 = op(e6, e2)) | (e6 = op(e5, e2)) | (e6 = op(e4, e2)) | (e6 = op(e3, e2)) | (e6 = op(e2, e2)) | (e6 = op(e1, e2)) | (e6 = op(e0, e2))) & ((e6 = op(e2, e6)) | (e6 = op(e2, e5)) | (e6 = op(e2, e4)) | (e6 = op(e2, e3)) | (e6 = op(e2, e2)) | (e6 = op(e2, e1)) | (e6 = op(e2, e0))) & ((e5 = op(e6, e2)) | (e5 = op(e5, e2)) | (e5 = op(e4, e2)) | (e5 = op(e3, e2)) | (e5 = op(e2, e2)) | (e5 = op(e1, e2)) | (e5 = op(e0, e2))) & ((e5 = op(e2, e6)) | (e5 = op(e2, e5)) | (e5 = op(e2, e4)) | (e5 = op(e2, e3)) | (e5 = op(e2, e2)) | (e5 = op(e2, e1)) | (e5 = op(e2, e0))) & ((e4 = op(e6, e2)) | (e4 = op(e5, e2)) | (e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e6)) | (e4 = op(e2, e5)) | (e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e6, e2)) | (e3 = op(e5, e2)) | (e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e6)) | (e3 = op(e2, e5)) | (e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e6, e2)) | (e2 = op(e5, e2)) | (e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e6)) | (e2 = op(e2, e5)) | (e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e6, e2)) | (e1 = op(e5, e2)) | (e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e6)) | (e1 = op(e2, e5)) | (e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e6, e2)) | (e0 = op(e5, e2)) | (e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e6)) | (e0 = op(e2, e5)) | (e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e6 = op(e6, e1)) | (e6 = op(e5, e1)) | (e6 = op(e4, e1)) | (e6 = op(e3, e1)) | (e6 = op(e2, e1)) | (e6 = op(e1, e1)) | (e6 = op(e0, e1))) & ((e6 = op(e1, e6)) | (e6 = op(e1, e5)) | (e6 = op(e1, e4)) | (e6 = op(e1, e3)) | (e6 = op(e1, e2)) | (e6 = op(e1, e1)) | (e6 = op(e1, e0))) & ((e5 = op(e6, e1)) | (e5 = op(e5, e1)) | (e5 = op(e4, e1)) | (e5 = op(e3, e1)) | (e5 = op(e2, e1)) | (e5 = op(e1, e1)) | (e5 = op(e0, e1))) & ((e5 = op(e1, e6)) | (e5 = op(e1, e5)) | (e5 = op(e1, e4)) | (e5 = op(e1, e3)) | (e5 = op(e1, e2)) | (e5 = op(e1, e1)) | (e5 = op(e1, e0))) & ((e4 = op(e6, e1)) | (e4 = op(e5, e1)) | (e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e6)) | (e4 = op(e1, e5)) | (e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e6, e1)) | (e3 = op(e5, e1)) | (e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e6)) | (e3 = op(e1, e5)) | (e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e6, e1)) | (e2 = op(e5, e1)) | (e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e6)) | (e2 = op(e1, e5)) | (e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e6, e1)) | (e1 = op(e5, e1)) | (e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e6)) | (e1 = op(e1, e5)) | (e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e6, e1)) | (e0 = op(e5, e1)) | (e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e6)) | (e0 = op(e1, e5)) | (e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e6 = op(e6, e0)) | (e6 = op(e5, e0)) | (e6 = op(e4, e0)) | (e6 = op(e3, e0)) | (e6 = op(e2, e0)) | (e6 = op(e1, e0)) | (e6 = op(e0, e0))) & ((e6 = op(e0, e6)) | (e6 = op(e0, e5)) | (e6 = op(e0, e4)) | (e6 = op(e0, e3)) | (e6 = op(e0, e2)) | (e6 = op(e0, e1)) | (e6 = op(e0, e0))) & ((e5 = op(e6, e0)) | (e5 = op(e5, e0)) | (e5 = op(e4, e0)) | (e5 = op(e3, e0)) | (e5 = op(e2, e0)) | (e5 = op(e1, e0)) | (e5 = op(e0, e0))) & ((e5 = op(e0, e6)) | (e5 = op(e0, e5)) | (e5 = op(e0, e4)) | (e5 = op(e0, e3)) | (e5 = op(e0, e2)) | (e5 = op(e0, e1)) | (e5 = op(e0, e0))) & ((e4 = op(e6, e0)) | (e4 = op(e5, e0)) | (e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e6)) | (e4 = op(e0, e5)) | (e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e6, e0)) | (e3 = op(e5, e0)) | (e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e6)) | (e3 = op(e0, e5)) | (e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e6, e0)) | (e2 = op(e5, e0)) | (e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e6)) | (e2 = op(e0, e5)) | (e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e6, e0)) | (e1 = op(e5, e0)) | (e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e6)) | (e1 = op(e0, e5)) | (e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e6, e0)) | (e0 = op(e5, e0)) | (e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e6)) | (e0 = op(e0, e5)) | (e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e6 = op(e6, e6)) | (e5 = op(e6, e6)) | (e4 = op(e6, e6)) | (e3 = op(e6, e6)) | (e2 = op(e6, e6)) | (e1 = op(e6, e6)) | (e0 = op(e6, e6))) & ((e6 = op(e6, e5)) | (e5 = op(e6, e5)) | (e4 = op(e6, e5)) | (e3 = op(e6, e5)) | (e2 = op(e6, e5)) | (e1 = op(e6, e5)) | (e0 = op(e6, e5))) & ((e6 = op(e6, e4)) | (e5 = op(e6, e4)) | (e4 = op(e6, e4)) | (e3 = op(e6, e4)) | (e2 = op(e6, e4)) | (e1 = op(e6, e4)) | (e0 = op(e6, e4))) & ((e6 = op(e6, e3)) | (e5 = op(e6, e3)) | (e4 = op(e6, e3)) | (e3 = op(e6, e3)) | (e2 = op(e6, e3)) | (e1 = op(e6, e3)) | (e0 = op(e6, e3))) & ((e6 = op(e6, e2)) | (e5 = op(e6, e2)) | (e4 = op(e6, e2)) | (e3 = op(e6, e2)) | (e2 = op(e6, e2)) | (e1 = op(e6, e2)) | (e0 = op(e6, e2))) & ((e6 = op(e6, e1)) | (e5 = op(e6, e1)) | (e4 = op(e6, e1)) | (e3 = op(e6, e1)) | (e2 = op(e6, e1)) | (e1 = op(e6, e1)) | (e0 = op(e6, e1))) & ((e6 = op(e6, e0)) | (e5 = op(e6, e0)) | (e4 = op(e6, e0)) | (e3 = op(e6, e0)) | (e2 = op(e6, e0)) | (e1 = op(e6, e0)) | (e0 = op(e6, e0))) & ((e6 = op(e5, e6)) | (e5 = op(e5, e6)) | (e4 = op(e5, e6)) | (e3 = op(e5, e6)) | (e2 = op(e5, e6)) | (e1 = op(e5, e6)) | (e0 = op(e5, e6))) & ((e6 = op(e5, e5)) | (e5 = op(e5, e5)) | (e4 = op(e5, e5)) | (e3 = op(e5, e5)) | (e2 = op(e5, e5)) | (e1 = op(e5, e5)) | (e0 = op(e5, e5))) & ((e6 = op(e5, e4)) | (e5 = op(e5, e4)) | (e4 = op(e5, e4)) | (e3 = op(e5, e4)) | (e2 = op(e5, e4)) | (e1 = op(e5, e4)) | (e0 = op(e5, e4))) & ((e6 = op(e5, e3)) | (e5 = op(e5, e3)) | (e4 = op(e5, e3)) | (e3 = op(e5, e3)) | (e2 = op(e5, e3)) | (e1 = op(e5, e3)) | (e0 = op(e5, e3))) & ((e6 = op(e5, e2)) | (e5 = op(e5, e2)) | (e4 = op(e5, e2)) | (e3 = op(e5, e2)) | (e2 = op(e5, e2)) | (e1 = op(e5, e2)) | (e0 = op(e5, e2))) & ((e6 = op(e5, e1)) | (e5 = op(e5, e1)) | (e4 = op(e5, e1)) | (e3 = op(e5, e1)) | (e2 = op(e5, e1)) | (e1 = op(e5, e1)) | (e0 = op(e5, e1))) & ((e6 = op(e5, e0)) | (e5 = op(e5, e0)) | (e4 = op(e5, e0)) | (e3 = op(e5, e0)) | (e2 = op(e5, e0)) | (e1 = op(e5, e0)) | (e0 = op(e5, e0))) & ((e6 = op(e4, e6)) | (e5 = op(e4, e6)) | (e4 = op(e4, e6)) | (e3 = op(e4, e6)) | (e2 = op(e4, e6)) | (e1 = op(e4, e6)) | (e0 = op(e4, e6))) & ((e6 = op(e4, e5)) | (e5 = op(e4, e5)) | (e4 = op(e4, e5)) | (e3 = op(e4, e5)) | (e2 = op(e4, e5)) | (e1 = op(e4, e5)) | (e0 = op(e4, e5))) & ((e6 = op(e4, e4)) | (e5 = op(e4, e4)) | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e6 = op(e4, e3)) | (e5 = op(e4, e3)) | (e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e6 = op(e4, e2)) | (e5 = op(e4, e2)) | (e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e6 = op(e4, e1)) | (e5 = op(e4, e1)) | (e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e6 = op(e4, e0)) | (e5 = op(e4, e0)) | (e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e6 = op(e3, e6)) | (e5 = op(e3, e6)) | (e4 = op(e3, e6)) | (e3 = op(e3, e6)) | (e2 = op(e3, e6)) | (e1 = op(e3, e6)) | (e0 = op(e3, e6))) & ((e6 = op(e3, e5)) | (e5 = op(e3, e5)) | (e4 = op(e3, e5)) | (e3 = op(e3, e5)) | (e2 = op(e3, e5)) | (e1 = op(e3, e5)) | (e0 = op(e3, e5))) & ((e6 = op(e3, e4)) | (e5 = op(e3, e4)) | (e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e6 = op(e3, e3)) | (e5 = op(e3, e3)) | (e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e6 = op(e3, e2)) | (e5 = op(e3, e2)) | (e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e6 = op(e3, e1)) | (e5 = op(e3, e1)) | (e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e6 = op(e3, e0)) | (e5 = op(e3, e0)) | (e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e6 = op(e2, e6)) | (e5 = op(e2, e6)) | (e4 = op(e2, e6)) | (e3 = op(e2, e6)) | (e2 = op(e2, e6)) | (e1 = op(e2, e6)) | (e0 = op(e2, e6))) & ((e6 = op(e2, e5)) | (e5 = op(e2, e5)) | (e4 = op(e2, e5)) | (e3 = op(e2, e5)) | (e2 = op(e2, e5)) | (e1 = op(e2, e5)) | (e0 = op(e2, e5))) & ((e6 = op(e2, e4)) | (e5 = op(e2, e4)) | (e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e6 = op(e2, e3)) | (e5 = op(e2, e3)) | (e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e6 = op(e2, e2)) | (e5 = op(e2, e2)) | (e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e6 = op(e2, e1)) | (e5 = op(e2, e1)) | (e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e6 = op(e2, e0)) | (e5 = op(e2, e0)) | (e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e6 = op(e1, e6)) | (e5 = op(e1, e6)) | (e4 = op(e1, e6)) | (e3 = op(e1, e6)) | (e2 = op(e1, e6)) | (e1 = op(e1, e6)) | (e0 = op(e1, e6))) & ((e6 = op(e1, e5)) | (e5 = op(e1, e5)) | (e4 = op(e1, e5)) | (e3 = op(e1, e5)) | (e2 = op(e1, e5)) | (e1 = op(e1, e5)) | (e0 = op(e1, e5))) & ((e6 = op(e1, e4)) | (e5 = op(e1, e4)) | (e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e6 = op(e1, e3)) | (e5 = op(e1, e3)) | (e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e6 = op(e1, e2)) | (e5 = op(e1, e2)) | (e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e6 = op(e1, e1)) | (e5 = op(e1, e1)) | (e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e6 = op(e1, e0)) | (e5 = op(e1, e0)) | (e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e6 = op(e0, e6)) | (e5 = op(e0, e6)) | (e4 = op(e0, e6)) | (e3 = op(e0, e6)) | (e2 = op(e0, e6)) | (e1 = op(e0, e6)) | (e0 = op(e0, e6))) & ((e6 = op(e0, e5)) | (e5 = op(e0, e5)) | (e4 = op(e0, e5)) | (e3 = op(e0, e5)) | (e2 = op(e0, e5)) | (e1 = op(e0, e5)) | (e0 = op(e0, e5))) & ((e6 = op(e0, e4)) | (e5 = op(e0, e4)) | (e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e6 = op(e0, e3)) | (e5 = op(e0, e3)) | (e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e6 = op(e0, e2)) | (e5 = op(e0, e2)) | (e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e6 = op(e0, e1)) | (e5 = op(e0, e1)) | (e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e6 = op(e0, e0)) | (e5 = op(e0, e0)) | (e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & (~ (e6 = op(e6, e6)) | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))) & ((e6 = op(e6, e6)) | (e5 = op(e5, e5)) | (e4 = op(e4, e4)) | (e3 = op(e3, e3)) | (e2 = op(e2, e2)) | (e1 = op(e1, e1)) | (e0 = op(e0, e0)))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ((e6 = op(op(op(e6, e6), e6), e6)) & (e6 = op(op(op(e5, e6), e5), e5)) & (e6 = op(op(op(e4, e6), e4), e4)) & (e6 = op(op(op(e3, e6), e3), e3)) & (e6 = op(op(op(e2, e6), e2), e2)) & (e6 = op(op(op(e1, e6), e1), e1)) & (e6 = op(op(op(e0, e6), e0), e0)) & (e5 = op(op(op(e6, e5), e6), e6)) & (e5 = op(op(op(e5, e5), e5), e5)) & (e5 = op(op(op(e4, e5), e4), e4)) & (e5 = op(op(op(e3, e5), e3), e3)) & (e5 = op(op(op(e2, e5), e2), e2)) & (e5 = op(op(op(e1, e5), e1), e1)) & (e5 = op(op(op(e0, e5), e0), e0)) & (e4 = op(op(op(e6, e4), e6), e6)) & (e4 = op(op(op(e5, e4), e5), e5)) & (e4 = op(op(op(e4, e4), e4), e4)) & (e4 = op(op(op(e3, e4), e3), e3)) & (e4 = op(op(op(e2, e4), e2), e2)) & (e4 = op(op(op(e1, e4), e1), e1)) & (e4 = op(op(op(e0, e4), e0), e0)) & (e3 = op(op(op(e6, e3), e6), e6)) & (e3 = op(op(op(e5, e3), e5), e5)) & (e3 = op(op(op(e4, e3), e4), e4)) & (e3 = op(op(op(e3, e3), e3), e3)) & (e3 = op(op(op(e2, e3), e2), e2)) & (e3 = op(op(op(e1, e3), e1), e1)) & (e3 = op(op(op(e0, e3), e0), e0)) & (e2 = op(op(op(e6, e2), e6), e6)) & (e2 = op(op(op(e5, e2), e5), e5)) & (e2 = op(op(op(e4, e2), e4), e4)) & (e2 = op(op(op(e3, e2), e3), e3)) & (e2 = op(op(op(e2, e2), e2), e2)) & (e2 = op(op(op(e1, e2), e1), e1)) & (e2 = op(op(op(e0, e2), e0), e0)) & (e1 = op(op(op(e6, e1), e6), e6)) & (e1 = op(op(op(e5, e1), e5), e5)) & (e1 = op(op(op(e4, e1), e4), e4)) & (e1 = op(op(op(e3, e1), e3), e3)) & (e1 = op(op(op(e2, e1), e2), e2)) & (e1 = op(op(op(e1, e1), e1), e1)) & (e1 = op(op(op(e0, e1), e0), e0)) & (e0 = op(op(op(e6, e0), e6), e6)) & (e0 = op(op(op(e5, e0), e5), e5)) & (e0 = op(op(op(e4, e0), e4), e4)) & (e0 = op(op(op(e3, e0), e3), e3)) & (e0 = op(op(op(e2, e0), e2), e2)) & (e0 = op(op(op(e1, e0), e1), e1)) & (e0 = op(op(op(e0, e0), e0), e0)) & ((e6 = op(e6, e6)) | (e6 = op(e5, e6)) | (e6 = op(e4, e6)) | (e6 = op(e3, e6)) | (e6 = op(e2, e6)) | (e6 = op(e1, e6)) | (e6 = op(e0, e6))) & ((e6 = op(e6, e6)) | (e6 = op(e6, e5)) | (e6 = op(e6, e4)) | (e6 = op(e6, e3)) | (e6 = op(e6, e2)) | (e6 = op(e6, e1)) | (e6 = op(e6, e0))) & ((e5 = op(e6, e6)) | (e5 = op(e5, e6)) | (e5 = op(e4, e6)) | (e5 = op(e3, e6)) | (e5 = op(e2, e6)) | (e5 = op(e1, e6)) | (e5 = op(e0, e6))) & ((e5 = op(e6, e6)) | (e5 = op(e6, e5)) | (e5 = op(e6, e4)) | (e5 = op(e6, e3)) | (e5 = op(e6, e2)) | (e5 = op(e6, e1)) | (e5 = op(e6, e0))) & ((e4 = op(e6, e6)) | (e4 = op(e5, e6)) | (e4 = op(e4, e6)) | (e4 = op(e3, e6)) | (e4 = op(e2, e6)) | (e4 = op(e1, e6)) | (e4 = op(e0, e6))) & ((e4 = op(e6, e6)) | (e4 = op(e6, e5)) | (e4 = op(e6, e4)) | (e4 = op(e6, e3)) | (e4 = op(e6, e2)) | (e4 = op(e6, e1)) | (e4 = op(e6, e0))) & ((e3 = op(e6, e6)) | (e3 = op(e5, e6)) | (e3 = op(e4, e6)) | (e3 = op(e3, e6)) | (e3 = op(e2, e6)) | (e3 = op(e1, e6)) | (e3 = op(e0, e6))) & ((e3 = op(e6, e6)) | (e3 = op(e6, e5)) | (e3 = op(e6, e4)) | (e3 = op(e6, e3)) | (e3 = op(e6, e2)) | (e3 = op(e6, e1)) | (e3 = op(e6, e0))) & ((e2 = op(e6, e6)) | (e2 = op(e5, e6)) | (e2 = op(e4, e6)) | (e2 = op(e3, e6)) | (e2 = op(e2, e6)) | (e2 = op(e1, e6)) | (e2 = op(e0, e6))) & ((e2 = op(e6, e6)) | (e2 = op(e6, e5)) | (e2 = op(e6, e4)) | (e2 = op(e6, e3)) | (e2 = op(e6, e2)) | (e2 = op(e6, e1)) | (e2 = op(e6, e0))) & ((e1 = op(e6, e6)) | (e1 = op(e5, e6)) | (e1 = op(e4, e6)) | (e1 = op(e3, e6)) | (e1 = op(e2, e6)) | (e1 = op(e1, e6)) | (e1 = op(e0, e6))) & ((e1 = op(e6, e6)) | (e1 = op(e6, e5)) | (e1 = op(e6, e4)) | (e1 = op(e6, e3)) | (e1 = op(e6, e2)) | (e1 = op(e6, e1)) | (e1 = op(e6, e0))) & ((e0 = op(e6, e6)) | (e0 = op(e5, e6)) | (e0 = op(e4, e6)) | (e0 = op(e3, e6)) | (e0 = op(e2, e6)) | (e0 = op(e1, e6)) | (e0 = op(e0, e6))) & ((e0 = op(e6, e6)) | (e0 = op(e6, e5)) | (e0 = op(e6, e4)) | (e0 = op(e6, e3)) | (e0 = op(e6, e2)) | (e0 = op(e6, e1)) | (e0 = op(e6, e0))) & ((e6 = op(e6, e5)) | (e6 = op(e5, e5)) | (e6 = op(e4, e5)) | (e6 = op(e3, e5)) | (e6 = op(e2, e5)) | (e6 = op(e1, e5)) | (e6 = op(e0, e5))) & ((e6 = op(e5, e6)) | (e6 = op(e5, e5)) | (e6 = op(e5, e4)) | (e6 = op(e5, e3)) | (e6 = op(e5, e2)) | (e6 = op(e5, e1)) | (e6 = op(e5, e0))) & ((e5 = op(e6, e5)) | (e5 = op(e5, e5)) | (e5 = op(e4, e5)) | (e5 = op(e3, e5)) | (e5 = op(e2, e5)) | (e5 = op(e1, e5)) | (e5 = op(e0, e5))) & ((e5 = op(e5, e6)) | (e5 = op(e5, e5)) | (e5 = op(e5, e4)) | (e5 = op(e5, e3)) | (e5 = op(e5, e2)) | (e5 = op(e5, e1)) | (e5 = op(e5, e0))) & ((e4 = op(e6, e5)) | (e4 = op(e5, e5)) | (e4 = op(e4, e5)) | (e4 = op(e3, e5)) | (e4 = op(e2, e5)) | (e4 = op(e1, e5)) | (e4 = op(e0, e5))) & ((e4 = op(e5, e6)) | (e4 = op(e5, e5)) | (e4 = op(e5, e4)) | (e4 = op(e5, e3)) | (e4 = op(e5, e2)) | (e4 = op(e5, e1)) | (e4 = op(e5, e0))) & ((e3 = op(e6, e5)) | (e3 = op(e5, e5)) | (e3 = op(e4, e5)) | (e3 = op(e3, e5)) | (e3 = op(e2, e5)) | (e3 = op(e1, e5)) | (e3 = op(e0, e5))) & ((e3 = op(e5, e6)) | (e3 = op(e5, e5)) | (e3 = op(e5, e4)) | (e3 = op(e5, e3)) | (e3 = op(e5, e2)) | (e3 = op(e5, e1)) | (e3 = op(e5, e0))) & ((e2 = op(e6, e5)) | (e2 = op(e5, e5)) | (e2 = op(e4, e5)) | (e2 = op(e3, e5)) | (e2 = op(e2, e5)) | (e2 = op(e1, e5)) | (e2 = op(e0, e5))) & ((e2 = op(e5, e6)) | (e2 = op(e5, e5)) | (e2 = op(e5, e4)) | (e2 = op(e5, e3)) | (e2 = op(e5, e2)) | (e2 = op(e5, e1)) | (e2 = op(e5, e0))) & ((e1 = op(e6, e5)) | (e1 = op(e5, e5)) | (e1 = op(e4, e5)) | (e1 = op(e3, e5)) | (e1 = op(e2, e5)) | (e1 = op(e1, e5)) | (e1 = op(e0, e5))) & ((e1 = op(e5, e6)) | (e1 = op(e5, e5)) | (e1 = op(e5, e4)) | (e1 = op(e5, e3)) | (e1 = op(e5, e2)) | (e1 = op(e5, e1)) | (e1 = op(e5, e0))) & ((e0 = op(e6, e5)) | (e0 = op(e5, e5)) | (e0 = op(e4, e5)) | (e0 = op(e3, e5)) | (e0 = op(e2, e5)) | (e0 = op(e1, e5)) | (e0 = op(e0, e5))) & ((e0 = op(e5, e6)) | (e0 = op(e5, e5)) | (e0 = op(e5, e4)) | (e0 = op(e5, e3)) | (e0 = op(e5, e2)) | (e0 = op(e5, e1)) | (e0 = op(e5, e0))) & ((e6 = op(e6, e4)) | (e6 = op(e5, e4)) | (e6 = op(e4, e4)) | (e6 = op(e3, e4)) | (e6 = op(e2, e4)) | (e6 = op(e1, e4)) | (e6 = op(e0, e4))) & ((e6 = op(e4, e6)) | (e6 = op(e4, e5)) | (e6 = op(e4, e4)) | (e6 = op(e4, e3)) | (e6 = op(e4, e2)) | (e6 = op(e4, e1)) | (e6 = op(e4, e0))) & ((e5 = op(e6, e4)) | (e5 = op(e5, e4)) | (e5 = op(e4, e4)) | (e5 = op(e3, e4)) | (e5 = op(e2, e4)) | (e5 = op(e1, e4)) | (e5 = op(e0, e4))) & ((e5 = op(e4, e6)) | (e5 = op(e4, e5)) | (e5 = op(e4, e4)) | (e5 = op(e4, e3)) | (e5 = op(e4, e2)) | (e5 = op(e4, e1)) | (e5 = op(e4, e0))) & ((e4 = op(e6, e4)) | (e4 = op(e5, e4)) | (e4 = op(e4, e4)) | (e4 = op(e3, e4)) | (e4 = op(e2, e4)) | (e4 = op(e1, e4)) | (e4 = op(e0, e4))) & ((e4 = op(e4, e6)) | (e4 = op(e4, e5)) | (e4 = op(e4, e4)) | (e4 = op(e4, e3)) | (e4 = op(e4, e2)) | (e4 = op(e4, e1)) | (e4 = op(e4, e0))) & ((e3 = op(e6, e4)) | (e3 = op(e5, e4)) | (e3 = op(e4, e4)) | (e3 = op(e3, e4)) | (e3 = op(e2, e4)) | (e3 = op(e1, e4)) | (e3 = op(e0, e4))) & ((e3 = op(e4, e6)) | (e3 = op(e4, e5)) | (e3 = op(e4, e4)) | (e3 = op(e4, e3)) | (e3 = op(e4, e2)) | (e3 = op(e4, e1)) | (e3 = op(e4, e0))) & ((e2 = op(e6, e4)) | (e2 = op(e5, e4)) | (e2 = op(e4, e4)) | (e2 = op(e3, e4)) | (e2 = op(e2, e4)) | (e2 = op(e1, e4)) | (e2 = op(e0, e4))) & ((e2 = op(e4, e6)) | (e2 = op(e4, e5)) | (e2 = op(e4, e4)) | (e2 = op(e4, e3)) | (e2 = op(e4, e2)) | (e2 = op(e4, e1)) | (e2 = op(e4, e0))) & ((e1 = op(e6, e4)) | (e1 = op(e5, e4)) | (e1 = op(e4, e4)) | (e1 = op(e3, e4)) | (e1 = op(e2, e4)) | (e1 = op(e1, e4)) | (e1 = op(e0, e4))) & ((e1 = op(e4, e6)) | (e1 = op(e4, e5)) | (e1 = op(e4, e4)) | (e1 = op(e4, e3)) | (e1 = op(e4, e2)) | (e1 = op(e4, e1)) | (e1 = op(e4, e0))) & ((e0 = op(e6, e4)) | (e0 = op(e5, e4)) | (e0 = op(e4, e4)) | (e0 = op(e3, e4)) | (e0 = op(e2, e4)) | (e0 = op(e1, e4)) | (e0 = op(e0, e4))) & ((e0 = op(e4, e6)) | (e0 = op(e4, e5)) | (e0 = op(e4, e4)) | (e0 = op(e4, e3)) | (e0 = op(e4, e2)) | (e0 = op(e4, e1)) | (e0 = op(e4, e0))) & ((e6 = op(e6, e3)) | (e6 = op(e5, e3)) | (e6 = op(e4, e3)) | (e6 = op(e3, e3)) | (e6 = op(e2, e3)) | (e6 = op(e1, e3)) | (e6 = op(e0, e3))) & ((e6 = op(e3, e6)) | (e6 = op(e3, e5)) | (e6 = op(e3, e4)) | (e6 = op(e3, e3)) | (e6 = op(e3, e2)) | (e6 = op(e3, e1)) | (e6 = op(e3, e0))) & ((e5 = op(e6, e3)) | (e5 = op(e5, e3)) | (e5 = op(e4, e3)) | (e5 = op(e3, e3)) | (e5 = op(e2, e3)) | (e5 = op(e1, e3)) | (e5 = op(e0, e3))) & ((e5 = op(e3, e6)) | (e5 = op(e3, e5)) | (e5 = op(e3, e4)) | (e5 = op(e3, e3)) | (e5 = op(e3, e2)) | (e5 = op(e3, e1)) | (e5 = op(e3, e0))) & ((e4 = op(e6, e3)) | (e4 = op(e5, e3)) | (e4 = op(e4, e3)) | (e4 = op(e3, e3)) | (e4 = op(e2, e3)) | (e4 = op(e1, e3)) | (e4 = op(e0, e3))) & ((e4 = op(e3, e6)) | (e4 = op(e3, e5)) | (e4 = op(e3, e4)) | (e4 = op(e3, e3)) | (e4 = op(e3, e2)) | (e4 = op(e3, e1)) | (e4 = op(e3, e0))) & ((e3 = op(e6, e3)) | (e3 = op(e5, e3)) | (e3 = op(e4, e3)) | (e3 = op(e3, e3)) | (e3 = op(e2, e3)) | (e3 = op(e1, e3)) | (e3 = op(e0, e3))) & ((e3 = op(e3, e6)) | (e3 = op(e3, e5)) | (e3 = op(e3, e4)) | (e3 = op(e3, e3)) | (e3 = op(e3, e2)) | (e3 = op(e3, e1)) | (e3 = op(e3, e0))) & ((e2 = op(e6, e3)) | (e2 = op(e5, e3)) | (e2 = op(e4, e3)) | (e2 = op(e3, e3)) | (e2 = op(e2, e3)) | (e2 = op(e1, e3)) | (e2 = op(e0, e3))) & ((e2 = op(e3, e6)) | (e2 = op(e3, e5)) | (e2 = op(e3, e4)) | (e2 = op(e3, e3)) | (e2 = op(e3, e2)) | (e2 = op(e3, e1)) | (e2 = op(e3, e0))) & ((e1 = op(e6, e3)) | (e1 = op(e5, e3)) | (e1 = op(e4, e3)) | (e1 = op(e3, e3)) | (e1 = op(e2, e3)) | (e1 = op(e1, e3)) | (e1 = op(e0, e3))) & ((e1 = op(e3, e6)) | (e1 = op(e3, e5)) | (e1 = op(e3, e4)) | (e1 = op(e3, e3)) | (e1 = op(e3, e2)) | (e1 = op(e3, e1)) | (e1 = op(e3, e0))) & ((e0 = op(e6, e3)) | (e0 = op(e5, e3)) | (e0 = op(e4, e3)) | (e0 = op(e3, e3)) | (e0 = op(e2, e3)) | (e0 = op(e1, e3)) | (e0 = op(e0, e3))) & ((e0 = op(e3, e6)) | (e0 = op(e3, e5)) | (e0 = op(e3, e4)) | (e0 = op(e3, e3)) | (e0 = op(e3, e2)) | (e0 = op(e3, e1)) | (e0 = op(e3, e0))) & ((e6 = op(e6, e2)) | (e6 = op(e5, e2)) | (e6 = op(e4, e2)) | (e6 = op(e3, e2)) | (e6 = op(e2, e2)) | (e6 = op(e1, e2)) | (e6 = op(e0, e2))) & ((e6 = op(e2, e6)) | (e6 = op(e2, e5)) | (e6 = op(e2, e4)) | (e6 = op(e2, e3)) | (e6 = op(e2, e2)) | (e6 = op(e2, e1)) | (e6 = op(e2, e0))) & ((e5 = op(e6, e2)) | (e5 = op(e5, e2)) | (e5 = op(e4, e2)) | (e5 = op(e3, e2)) | (e5 = op(e2, e2)) | (e5 = op(e1, e2)) | (e5 = op(e0, e2))) & ((e5 = op(e2, e6)) | (e5 = op(e2, e5)) | (e5 = op(e2, e4)) | (e5 = op(e2, e3)) | (e5 = op(e2, e2)) | (e5 = op(e2, e1)) | (e5 = op(e2, e0))) & ((e4 = op(e6, e2)) | (e4 = op(e5, e2)) | (e4 = op(e4, e2)) | (e4 = op(e3, e2)) | (e4 = op(e2, e2)) | (e4 = op(e1, e2)) | (e4 = op(e0, e2))) & ((e4 = op(e2, e6)) | (e4 = op(e2, e5)) | (e4 = op(e2, e4)) | (e4 = op(e2, e3)) | (e4 = op(e2, e2)) | (e4 = op(e2, e1)) | (e4 = op(e2, e0))) & ((e3 = op(e6, e2)) | (e3 = op(e5, e2)) | (e3 = op(e4, e2)) | (e3 = op(e3, e2)) | (e3 = op(e2, e2)) | (e3 = op(e1, e2)) | (e3 = op(e0, e2))) & ((e3 = op(e2, e6)) | (e3 = op(e2, e5)) | (e3 = op(e2, e4)) | (e3 = op(e2, e3)) | (e3 = op(e2, e2)) | (e3 = op(e2, e1)) | (e3 = op(e2, e0))) & ((e2 = op(e6, e2)) | (e2 = op(e5, e2)) | (e2 = op(e4, e2)) | (e2 = op(e3, e2)) | (e2 = op(e2, e2)) | (e2 = op(e1, e2)) | (e2 = op(e0, e2))) & ((e2 = op(e2, e6)) | (e2 = op(e2, e5)) | (e2 = op(e2, e4)) | (e2 = op(e2, e3)) | (e2 = op(e2, e2)) | (e2 = op(e2, e1)) | (e2 = op(e2, e0))) & ((e1 = op(e6, e2)) | (e1 = op(e5, e2)) | (e1 = op(e4, e2)) | (e1 = op(e3, e2)) | (e1 = op(e2, e2)) | (e1 = op(e1, e2)) | (e1 = op(e0, e2))) & ((e1 = op(e2, e6)) | (e1 = op(e2, e5)) | (e1 = op(e2, e4)) | (e1 = op(e2, e3)) | (e1 = op(e2, e2)) | (e1 = op(e2, e1)) | (e1 = op(e2, e0))) & ((e0 = op(e6, e2)) | (e0 = op(e5, e2)) | (e0 = op(e4, e2)) | (e0 = op(e3, e2)) | (e0 = op(e2, e2)) | (e0 = op(e1, e2)) | (e0 = op(e0, e2))) & ((e0 = op(e2, e6)) | (e0 = op(e2, e5)) | (e0 = op(e2, e4)) | (e0 = op(e2, e3)) | (e0 = op(e2, e2)) | (e0 = op(e2, e1)) | (e0 = op(e2, e0))) & ((e6 = op(e6, e1)) | (e6 = op(e5, e1)) | (e6 = op(e4, e1)) | (e6 = op(e3, e1)) | (e6 = op(e2, e1)) | (e6 = op(e1, e1)) | (e6 = op(e0, e1))) & ((e6 = op(e1, e6)) | (e6 = op(e1, e5)) | (e6 = op(e1, e4)) | (e6 = op(e1, e3)) | (e6 = op(e1, e2)) | (e6 = op(e1, e1)) | (e6 = op(e1, e0))) & ((e5 = op(e6, e1)) | (e5 = op(e5, e1)) | (e5 = op(e4, e1)) | (e5 = op(e3, e1)) | (e5 = op(e2, e1)) | (e5 = op(e1, e1)) | (e5 = op(e0, e1))) & ((e5 = op(e1, e6)) | (e5 = op(e1, e5)) | (e5 = op(e1, e4)) | (e5 = op(e1, e3)) | (e5 = op(e1, e2)) | (e5 = op(e1, e1)) | (e5 = op(e1, e0))) & ((e4 = op(e6, e1)) | (e4 = op(e5, e1)) | (e4 = op(e4, e1)) | (e4 = op(e3, e1)) | (e4 = op(e2, e1)) | (e4 = op(e1, e1)) | (e4 = op(e0, e1))) & ((e4 = op(e1, e6)) | (e4 = op(e1, e5)) | (e4 = op(e1, e4)) | (e4 = op(e1, e3)) | (e4 = op(e1, e2)) | (e4 = op(e1, e1)) | (e4 = op(e1, e0))) & ((e3 = op(e6, e1)) | (e3 = op(e5, e1)) | (e3 = op(e4, e1)) | (e3 = op(e3, e1)) | (e3 = op(e2, e1)) | (e3 = op(e1, e1)) | (e3 = op(e0, e1))) & ((e3 = op(e1, e6)) | (e3 = op(e1, e5)) | (e3 = op(e1, e4)) | (e3 = op(e1, e3)) | (e3 = op(e1, e2)) | (e3 = op(e1, e1)) | (e3 = op(e1, e0))) & ((e2 = op(e6, e1)) | (e2 = op(e5, e1)) | (e2 = op(e4, e1)) | (e2 = op(e3, e1)) | (e2 = op(e2, e1)) | (e2 = op(e1, e1)) | (e2 = op(e0, e1))) & ((e2 = op(e1, e6)) | (e2 = op(e1, e5)) | (e2 = op(e1, e4)) | (e2 = op(e1, e3)) | (e2 = op(e1, e2)) | (e2 = op(e1, e1)) | (e2 = op(e1, e0))) & ((e1 = op(e6, e1)) | (e1 = op(e5, e1)) | (e1 = op(e4, e1)) | (e1 = op(e3, e1)) | (e1 = op(e2, e1)) | (e1 = op(e1, e1)) | (e1 = op(e0, e1))) & ((e1 = op(e1, e6)) | (e1 = op(e1, e5)) | (e1 = op(e1, e4)) | (e1 = op(e1, e3)) | (e1 = op(e1, e2)) | (e1 = op(e1, e1)) | (e1 = op(e1, e0))) & ((e0 = op(e6, e1)) | (e0 = op(e5, e1)) | (e0 = op(e4, e1)) | (e0 = op(e3, e1)) | (e0 = op(e2, e1)) | (e0 = op(e1, e1)) | (e0 = op(e0, e1))) & ((e0 = op(e1, e6)) | (e0 = op(e1, e5)) | (e0 = op(e1, e4)) | (e0 = op(e1, e3)) | (e0 = op(e1, e2)) | (e0 = op(e1, e1)) | (e0 = op(e1, e0))) & ((e6 = op(e6, e0)) | (e6 = op(e5, e0)) | (e6 = op(e4, e0)) | (e6 = op(e3, e0)) | (e6 = op(e2, e0)) | (e6 = op(e1, e0)) | (e6 = op(e0, e0))) & ((e6 = op(e0, e6)) | (e6 = op(e0, e5)) | (e6 = op(e0, e4)) | (e6 = op(e0, e3)) | (e6 = op(e0, e2)) | (e6 = op(e0, e1)) | (e6 = op(e0, e0))) & ((e5 = op(e6, e0)) | (e5 = op(e5, e0)) | (e5 = op(e4, e0)) | (e5 = op(e3, e0)) | (e5 = op(e2, e0)) | (e5 = op(e1, e0)) | (e5 = op(e0, e0))) & ((e5 = op(e0, e6)) | (e5 = op(e0, e5)) | (e5 = op(e0, e4)) | (e5 = op(e0, e3)) | (e5 = op(e0, e2)) | (e5 = op(e0, e1)) | (e5 = op(e0, e0))) & ((e4 = op(e6, e0)) | (e4 = op(e5, e0)) | (e4 = op(e4, e0)) | (e4 = op(e3, e0)) | (e4 = op(e2, e0)) | (e4 = op(e1, e0)) | (e4 = op(e0, e0))) & ((e4 = op(e0, e6)) | (e4 = op(e0, e5)) | (e4 = op(e0, e4)) | (e4 = op(e0, e3)) | (e4 = op(e0, e2)) | (e4 = op(e0, e1)) | (e4 = op(e0, e0))) & ((e3 = op(e6, e0)) | (e3 = op(e5, e0)) | (e3 = op(e4, e0)) | (e3 = op(e3, e0)) | (e3 = op(e2, e0)) | (e3 = op(e1, e0)) | (e3 = op(e0, e0))) & ((e3 = op(e0, e6)) | (e3 = op(e0, e5)) | (e3 = op(e0, e4)) | (e3 = op(e0, e3)) | (e3 = op(e0, e2)) | (e3 = op(e0, e1)) | (e3 = op(e0, e0))) & ((e2 = op(e6, e0)) | (e2 = op(e5, e0)) | (e2 = op(e4, e0)) | (e2 = op(e3, e0)) | (e2 = op(e2, e0)) | (e2 = op(e1, e0)) | (e2 = op(e0, e0))) & ((e2 = op(e0, e6)) | (e2 = op(e0, e5)) | (e2 = op(e0, e4)) | (e2 = op(e0, e3)) | (e2 = op(e0, e2)) | (e2 = op(e0, e1)) | (e2 = op(e0, e0))) & ((e1 = op(e6, e0)) | (e1 = op(e5, e0)) | (e1 = op(e4, e0)) | (e1 = op(e3, e0)) | (e1 = op(e2, e0)) | (e1 = op(e1, e0)) | (e1 = op(e0, e0))) & ((e1 = op(e0, e6)) | (e1 = op(e0, e5)) | (e1 = op(e0, e4)) | (e1 = op(e0, e3)) | (e1 = op(e0, e2)) | (e1 = op(e0, e1)) | (e1 = op(e0, e0))) & ((e0 = op(e6, e0)) | (e0 = op(e5, e0)) | (e0 = op(e4, e0)) | (e0 = op(e3, e0)) | (e0 = op(e2, e0)) | (e0 = op(e1, e0)) | (e0 = op(e0, e0))) & ((e0 = op(e0, e6)) | (e0 = op(e0, e5)) | (e0 = op(e0, e4)) | (e0 = op(e0, e3)) | (e0 = op(e0, e2)) | (e0 = op(e0, e1)) | (e0 = op(e0, e0))) & ((e6 = op(e6, e6)) | (e5 = op(e6, e6)) | (e4 = op(e6, e6)) | (e3 = op(e6, e6)) | (e2 = op(e6, e6)) | (e1 = op(e6, e6)) | (e0 = op(e6, e6))) & ((e6 = op(e6, e5)) | (e5 = op(e6, e5)) | (e4 = op(e6, e5)) | (e3 = op(e6, e5)) | (e2 = op(e6, e5)) | (e1 = op(e6, e5)) | (e0 = op(e6, e5))) & ((e6 = op(e6, e4)) | (e5 = op(e6, e4)) | (e4 = op(e6, e4)) | (e3 = op(e6, e4)) | (e2 = op(e6, e4)) | (e1 = op(e6, e4)) | (e0 = op(e6, e4))) & ((e6 = op(e6, e3)) | (e5 = op(e6, e3)) | (e4 = op(e6, e3)) | (e3 = op(e6, e3)) | (e2 = op(e6, e3)) | (e1 = op(e6, e3)) | (e0 = op(e6, e3))) & ((e6 = op(e6, e2)) | (e5 = op(e6, e2)) | (e4 = op(e6, e2)) | (e3 = op(e6, e2)) | (e2 = op(e6, e2)) | (e1 = op(e6, e2)) | (e0 = op(e6, e2))) & ((e6 = op(e6, e1)) | (e5 = op(e6, e1)) | (e4 = op(e6, e1)) | (e3 = op(e6, e1)) | (e2 = op(e6, e1)) | (e1 = op(e6, e1)) | (e0 = op(e6, e1))) & ((e6 = op(e6, e0)) | (e5 = op(e6, e0)) | (e4 = op(e6, e0)) | (e3 = op(e6, e0)) | (e2 = op(e6, e0)) | (e1 = op(e6, e0)) | (e0 = op(e6, e0))) & ((e6 = op(e5, e6)) | (e5 = op(e5, e6)) | (e4 = op(e5, e6)) | (e3 = op(e5, e6)) | (e2 = op(e5, e6)) | (e1 = op(e5, e6)) | (e0 = op(e5, e6))) & ((e6 = op(e5, e5)) | (e5 = op(e5, e5)) | (e4 = op(e5, e5)) | (e3 = op(e5, e5)) | (e2 = op(e5, e5)) | (e1 = op(e5, e5)) | (e0 = op(e5, e5))) & ((e6 = op(e5, e4)) | (e5 = op(e5, e4)) | (e4 = op(e5, e4)) | (e3 = op(e5, e4)) | (e2 = op(e5, e4)) | (e1 = op(e5, e4)) | (e0 = op(e5, e4))) & ((e6 = op(e5, e3)) | (e5 = op(e5, e3)) | (e4 = op(e5, e3)) | (e3 = op(e5, e3)) | (e2 = op(e5, e3)) | (e1 = op(e5, e3)) | (e0 = op(e5, e3))) & ((e6 = op(e5, e2)) | (e5 = op(e5, e2)) | (e4 = op(e5, e2)) | (e3 = op(e5, e2)) | (e2 = op(e5, e2)) | (e1 = op(e5, e2)) | (e0 = op(e5, e2))) & ((e6 = op(e5, e1)) | (e5 = op(e5, e1)) | (e4 = op(e5, e1)) | (e3 = op(e5, e1)) | (e2 = op(e5, e1)) | (e1 = op(e5, e1)) | (e0 = op(e5, e1))) & ((e6 = op(e5, e0)) | (e5 = op(e5, e0)) | (e4 = op(e5, e0)) | (e3 = op(e5, e0)) | (e2 = op(e5, e0)) | (e1 = op(e5, e0)) | (e0 = op(e5, e0))) & ((e6 = op(e4, e6)) | (e5 = op(e4, e6)) | (e4 = op(e4, e6)) | (e3 = op(e4, e6)) | (e2 = op(e4, e6)) | (e1 = op(e4, e6)) | (e0 = op(e4, e6))) & ((e6 = op(e4, e5)) | (e5 = op(e4, e5)) | (e4 = op(e4, e5)) | (e3 = op(e4, e5)) | (e2 = op(e4, e5)) | (e1 = op(e4, e5)) | (e0 = op(e4, e5))) & ((e6 = op(e4, e4)) | (e5 = op(e4, e4)) | (e4 = op(e4, e4)) | (e3 = op(e4, e4)) | (e2 = op(e4, e4)) | (e1 = op(e4, e4)) | (e0 = op(e4, e4))) & ((e6 = op(e4, e3)) | (e5 = op(e4, e3)) | (e4 = op(e4, e3)) | (e3 = op(e4, e3)) | (e2 = op(e4, e3)) | (e1 = op(e4, e3)) | (e0 = op(e4, e3))) & ((e6 = op(e4, e2)) | (e5 = op(e4, e2)) | (e4 = op(e4, e2)) | (e3 = op(e4, e2)) | (e2 = op(e4, e2)) | (e1 = op(e4, e2)) | (e0 = op(e4, e2))) & ((e6 = op(e4, e1)) | (e5 = op(e4, e1)) | (e4 = op(e4, e1)) | (e3 = op(e4, e1)) | (e2 = op(e4, e1)) | (e1 = op(e4, e1)) | (e0 = op(e4, e1))) & ((e6 = op(e4, e0)) | (e5 = op(e4, e0)) | (e4 = op(e4, e0)) | (e3 = op(e4, e0)) | (e2 = op(e4, e0)) | (e1 = op(e4, e0)) | (e0 = op(e4, e0))) & ((e6 = op(e3, e6)) | (e5 = op(e3, e6)) | (e4 = op(e3, e6)) | (e3 = op(e3, e6)) | (e2 = op(e3, e6)) | (e1 = op(e3, e6)) | (e0 = op(e3, e6))) & ((e6 = op(e3, e5)) | (e5 = op(e3, e5)) | (e4 = op(e3, e5)) | (e3 = op(e3, e5)) | (e2 = op(e3, e5)) | (e1 = op(e3, e5)) | (e0 = op(e3, e5))) & ((e6 = op(e3, e4)) | (e5 = op(e3, e4)) | (e4 = op(e3, e4)) | (e3 = op(e3, e4)) | (e2 = op(e3, e4)) | (e1 = op(e3, e4)) | (e0 = op(e3, e4))) & ((e6 = op(e3, e3)) | (e5 = op(e3, e3)) | (e4 = op(e3, e3)) | (e3 = op(e3, e3)) | (e2 = op(e3, e3)) | (e1 = op(e3, e3)) | (e0 = op(e3, e3))) & ((e6 = op(e3, e2)) | (e5 = op(e3, e2)) | (e4 = op(e3, e2)) | (e3 = op(e3, e2)) | (e2 = op(e3, e2)) | (e1 = op(e3, e2)) | (e0 = op(e3, e2))) & ((e6 = op(e3, e1)) | (e5 = op(e3, e1)) | (e4 = op(e3, e1)) | (e3 = op(e3, e1)) | (e2 = op(e3, e1)) | (e1 = op(e3, e1)) | (e0 = op(e3, e1))) & ((e6 = op(e3, e0)) | (e5 = op(e3, e0)) | (e4 = op(e3, e0)) | (e3 = op(e3, e0)) | (e2 = op(e3, e0)) | (e1 = op(e3, e0)) | (e0 = op(e3, e0))) & ((e6 = op(e2, e6)) | (e5 = op(e2, e6)) | (e4 = op(e2, e6)) | (e3 = op(e2, e6)) | (e2 = op(e2, e6)) | (e1 = op(e2, e6)) | (e0 = op(e2, e6))) & ((e6 = op(e2, e5)) | (e5 = op(e2, e5)) | (e4 = op(e2, e5)) | (e3 = op(e2, e5)) | (e2 = op(e2, e5)) | (e1 = op(e2, e5)) | (e0 = op(e2, e5))) & ((e6 = op(e2, e4)) | (e5 = op(e2, e4)) | (e4 = op(e2, e4)) | (e3 = op(e2, e4)) | (e2 = op(e2, e4)) | (e1 = op(e2, e4)) | (e0 = op(e2, e4))) & ((e6 = op(e2, e3)) | (e5 = op(e2, e3)) | (e4 = op(e2, e3)) | (e3 = op(e2, e3)) | (e2 = op(e2, e3)) | (e1 = op(e2, e3)) | (e0 = op(e2, e3))) & ((e6 = op(e2, e2)) | (e5 = op(e2, e2)) | (e4 = op(e2, e2)) | (e3 = op(e2, e2)) | (e2 = op(e2, e2)) | (e1 = op(e2, e2)) | (e0 = op(e2, e2))) & ((e6 = op(e2, e1)) | (e5 = op(e2, e1)) | (e4 = op(e2, e1)) | (e3 = op(e2, e1)) | (e2 = op(e2, e1)) | (e1 = op(e2, e1)) | (e0 = op(e2, e1))) & ((e6 = op(e2, e0)) | (e5 = op(e2, e0)) | (e4 = op(e2, e0)) | (e3 = op(e2, e0)) | (e2 = op(e2, e0)) | (e1 = op(e2, e0)) | (e0 = op(e2, e0))) & ((e6 = op(e1, e6)) | (e5 = op(e1, e6)) | (e4 = op(e1, e6)) | (e3 = op(e1, e6)) | (e2 = op(e1, e6)) | (e1 = op(e1, e6)) | (e0 = op(e1, e6))) & ((e6 = op(e1, e5)) | (e5 = op(e1, e5)) | (e4 = op(e1, e5)) | (e3 = op(e1, e5)) | (e2 = op(e1, e5)) | (e1 = op(e1, e5)) | (e0 = op(e1, e5))) & ((e6 = op(e1, e4)) | (e5 = op(e1, e4)) | (e4 = op(e1, e4)) | (e3 = op(e1, e4)) | (e2 = op(e1, e4)) | (e1 = op(e1, e4)) | (e0 = op(e1, e4))) & ((e6 = op(e1, e3)) | (e5 = op(e1, e3)) | (e4 = op(e1, e3)) | (e3 = op(e1, e3)) | (e2 = op(e1, e3)) | (e1 = op(e1, e3)) | (e0 = op(e1, e3))) & ((e6 = op(e1, e2)) | (e5 = op(e1, e2)) | (e4 = op(e1, e2)) | (e3 = op(e1, e2)) | (e2 = op(e1, e2)) | (e1 = op(e1, e2)) | (e0 = op(e1, e2))) & ((e6 = op(e1, e1)) | (e5 = op(e1, e1)) | (e4 = op(e1, e1)) | (e3 = op(e1, e1)) | (e2 = op(e1, e1)) | (e1 = op(e1, e1)) | (e0 = op(e1, e1))) & ((e6 = op(e1, e0)) | (e5 = op(e1, e0)) | (e4 = op(e1, e0)) | (e3 = op(e1, e0)) | (e2 = op(e1, e0)) | (e1 = op(e1, e0)) | (e0 = op(e1, e0))) & ((e6 = op(e0, e6)) | (e5 = op(e0, e6)) | (e4 = op(e0, e6)) | (e3 = op(e0, e6)) | (e2 = op(e0, e6)) | (e1 = op(e0, e6)) | (e0 = op(e0, e6))) & ((e6 = op(e0, e5)) | (e5 = op(e0, e5)) | (e4 = op(e0, e5)) | (e3 = op(e0, e5)) | (e2 = op(e0, e5)) | (e1 = op(e0, e5)) | (e0 = op(e0, e5))) & ((e6 = op(e0, e4)) | (e5 = op(e0, e4)) | (e4 = op(e0, e4)) | (e3 = op(e0, e4)) | (e2 = op(e0, e4)) | (e1 = op(e0, e4)) | (e0 = op(e0, e4))) & ((e6 = op(e0, e3)) | (e5 = op(e0, e3)) | (e4 = op(e0, e3)) | (e3 = op(e0, e3)) | (e2 = op(e0, e3)) | (e1 = op(e0, e3)) | (e0 = op(e0, e3))) & ((e6 = op(e0, e2)) | (e5 = op(e0, e2)) | (e4 = op(e0, e2)) | (e3 = op(e0, e2)) | (e2 = op(e0, e2)) | (e1 = op(e0, e2)) | (e0 = op(e0, e2))) & ((e6 = op(e0, e1)) | (e5 = op(e0, e1)) | (e4 = op(e0, e1)) | (e3 = op(e0, e1)) | (e2 = op(e0, e1)) | (e1 = op(e0, e1)) | (e0 = op(e0, e1))) & ((e6 = op(e0, e0)) | (e5 = op(e0, e0)) | (e4 = op(e0, e0)) | (e3 = op(e0, e0)) | (e2 = op(e0, e0)) | (e1 = op(e0, e0)) | (e0 = op(e0, e0))) & (~ (e6 = op(e6, e6)) | ~ (e5 = op(e5, e5)) | ~ (e4 = op(e4, e4)) | ~ (e3 = op(e3, e3)) | ~ (e2 = op(e2, e2)) | ~ (e1 = op(e1, e1)) | ~ (e0 = op(e0, e0))) & ((e6 = op(e6, e6)) | (e5 = op(e5, e5)) | (e4 = op(e4, e4)) | (e3 = op(e3, e3)) | (e2 = op(e2, e2)) | (e1 = op(e1, e1)) | (e0 = op(e0, e0)))), file('/home/ubuntu/library/tptp/Problems/ALG/ALG209+1.p', co1)).