fof(f38488, plain, $false, inference(avatar_sat_refutation, [], [f600, f610, f620, f2150, f2916, f31991, f36378, f36625, f38487])).
fof(f38487, plain, (~ spl42_6 | ~ spl42_8 | spl42_1048), inference(avatar_contradiction_clause, [], [f38486])).
fof(f38486, plain, ($false | (~ spl42_6 | ~ spl42_8 | spl42_1048)), inference(subsumption_resolution, [], [f38485, f609])).
fof(f609, plain, (l1_cat_1(k1_yoneda_1(sK6)) | ~ spl42_8), inference(avatar_component_clause, [], [f607])).
fof(f607, plain, (spl42_8 <=> l1_cat_1(k1_yoneda_1(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl42_8])])).
fof(f38485, plain, (~ l1_cat_1(k1_yoneda_1(sK6)) | (~ spl42_6 | ~ spl42_8 | spl42_1048)), inference(subsumption_resolution, [], [f38484, f294])).
fof(f294, plain, v2_cat_1(sK6), inference(cnf_transformation, [], [f226])).
fof(f226, plain, ((~ m1_subset_1(k4_tarski(k4_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7))), k3_yoneda_1(sK6, sK7)), u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))) & m1_subset_1(sK7, u2_cat_1(sK6))) & l1_cat_1(sK6) & v2_cat_1(sK6)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6, sK7])], [f104, f225, f224])).
fof(f224, plain, (? [X0] : (? [X1] : (~ m1_subset_1(k4_tarski(k4_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))), k3_yoneda_1(X0, X1)), u2_cat_1(k12_nattra_1(X0, k1_yoneda_1(X0)))) & m1_subset_1(X1, u2_cat_1(X0))) & l1_cat_1(X0) & v2_cat_1(X0)) => (? [X1] : (~ m1_subset_1(k4_tarski(k4_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, X1)), k2_yoneda_1(sK6, k2_cat_1(sK6, X1))), k3_yoneda_1(sK6, X1)), u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))) & m1_subset_1(X1, u2_cat_1(sK6))) & l1_cat_1(sK6) & v2_cat_1(sK6))), introduced(choice_axiom, [])).
fof(f225, plain, (? [X1] : (~ m1_subset_1(k4_tarski(k4_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, X1)), k2_yoneda_1(sK6, k2_cat_1(sK6, X1))), k3_yoneda_1(sK6, X1)), u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))) & m1_subset_1(X1, u2_cat_1(sK6))) => (~ m1_subset_1(k4_tarski(k4_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7))), k3_yoneda_1(sK6, sK7)), u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))) & m1_subset_1(sK7, u2_cat_1(sK6)))), introduced(choice_axiom, [])).
fof(f104, plain, ? [X0] : (? [X1] : (~ m1_subset_1(k4_tarski(k4_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))), k3_yoneda_1(X0, X1)), u2_cat_1(k12_nattra_1(X0, k1_yoneda_1(X0)))) & m1_subset_1(X1, u2_cat_1(X0))) & l1_cat_1(X0) & v2_cat_1(X0)), inference(flattening, [], [f103])).
fof(f103, plain, ? [X0] : (? [X1] : (~ m1_subset_1(k4_tarski(k4_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))), k3_yoneda_1(X0, X1)), u2_cat_1(k12_nattra_1(X0, k1_yoneda_1(X0)))) & m1_subset_1(X1, u2_cat_1(X0))) & (l1_cat_1(X0) & v2_cat_1(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ! [X1] : (m1_subset_1(X1, u2_cat_1(X0)) => m1_subset_1(k4_tarski(k4_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))), k3_yoneda_1(X0, X1)), u2_cat_1(k12_nattra_1(X0, k1_yoneda_1(X0)))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ! [X1] : (m1_subset_1(X1, u2_cat_1(X0)) => m1_subset_1(k4_tarski(k4_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))), k3_yoneda_1(X0, X1)), u2_cat_1(k12_nattra_1(X0, k1_yoneda_1(X0)))))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', t4_yoneda_1)).
fof(f38484, plain, (~ v2_cat_1(sK6) | ~ l1_cat_1(k1_yoneda_1(sK6)) | (~ spl42_6 | ~ spl42_8 | spl42_1048)), inference(subsumption_resolution, [], [f38483, f295])).
fof(f295, plain, l1_cat_1(sK6), inference(cnf_transformation, [], [f226])).
fof(f38483, plain, (~ l1_cat_1(sK6) | ~ v2_cat_1(sK6) | ~ l1_cat_1(k1_yoneda_1(sK6)) | (~ spl42_6 | ~ spl42_8 | spl42_1048)), inference(subsumption_resolution, [], [f38479, f599])).
fof(f599, plain, (v2_cat_1(k1_yoneda_1(sK6)) | ~ spl42_6), inference(avatar_component_clause, [], [f597])).
fof(f597, plain, (spl42_6 <=> v2_cat_1(k1_yoneda_1(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl42_6])])).
fof(f38479, plain, (~ v2_cat_1(k1_yoneda_1(sK6)) | ~ l1_cat_1(sK6) | ~ v2_cat_1(sK6) | ~ l1_cat_1(k1_yoneda_1(sK6)) | (~ spl42_6 | ~ spl42_8 | spl42_1048)), inference(resolution, [], [f38404, f1633])).
fof(f1633, plain, ! [X0, X1] : (sP0(X0, X1, k11_nattra_1(X1, X0)) | ~ v2_cat_1(X0) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0)), inference(resolution, [], [f800, f469])).
fof(f469, plain, ! [X2, X1] : (~ sP1(k11_nattra_1(X1, X2), X1, X2) | sP0(X2, X1, k11_nattra_1(X1, X2))), inference(equality_resolution, [], [f305])).
fof(f305, plain, ! [X2, X0, X1] : (sP0(X2, X1, X0) | ~ (k11_nattra_1(X1, X2) = X0) | ~ sP1(X0, X1, X2)), inference(cnf_transformation, [], [f228])).
fof(f228, plain, ! [X0, X1, X2] : ((((k11_nattra_1(X1, X2) = X0) | ~ sP0(X2, X1, X0)) & (sP0(X2, X1, X0) | ~ (k11_nattra_1(X1, X2) = X0))) | ~ sP1(X0, X1, X2)), inference(rectify, [], [f227])).
fof(f227, plain, ! [X2, X0, X1] : ((((k11_nattra_1(X0, X1) = X2) | ~ sP0(X1, X0, X2)) & (sP0(X1, X0, X2) | ~ (k11_nattra_1(X0, X1) = X2))) | ~ sP1(X2, X0, X1)), inference(nnf_transformation, [], [f217])).
fof(f217, plain, ! [X2, X0, X1] : (((k11_nattra_1(X0, X1) = X2) <=> sP0(X1, X0, X2)) | ~ sP1(X2, X0, X1)), inference(usedef, [], [e217])).
fof(e217, plain, ! [X2, X0, X1] : (sP1(X2, X0, X1) <=> ((k11_nattra_1(X0, X1) = X2) <=> sP0(X1, X0, X2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f800, plain, ! [X0, X1] : (sP1(k11_nattra_1(X0, X1), X0, X1) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(duplicate_literal_removal, [], [f797])).
fof(f797, plain, ! [X0, X1] : (sP1(k11_nattra_1(X0, X1), X0, X1) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(resolution, [], [f319, f363])).
fof(f363, plain, ! [X0, X1] : (m4_nattra_1(k11_nattra_1(X0, X1), X0, X1) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ! [X0, X1] : (m4_nattra_1(k11_nattra_1(X0, X1), X0, X1) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(flattening, [], [f126])).
fof(f126, plain, ! [X0, X1] : (m4_nattra_1(k11_nattra_1(X0, X1), X0, X1) | (~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0, X1] : ((l1_cat_1(X1) & v2_cat_1(X1) & l1_cat_1(X0) & v2_cat_1(X0)) => m4_nattra_1(k11_nattra_1(X0, X1), X0, X1)), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', dt_k11_nattra_1)).
fof(f319, plain, ! [X2, X0, X1] : (~ m4_nattra_1(X2, X0, X1) | sP1(X2, X0, X1) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f218])).
fof(f218, plain, ! [X0] : (! [X1] : (! [X2] : (sP1(X2, X0, X1) | ~ m4_nattra_1(X2, X0, X1)) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(definition_folding, [], [f113, e217, e216])).
fof(f216, plain, ! [X1, X0, X2] : (sP0(X1, X0, X2) <=> ! [X3] : (r2_hidden(X3, X2) <=> ? [X4] : (? [X5] : (? [X6] : (r2_nattra_1(X0, X1, X4, X5) & (k4_tarski(k4_tarski(X4, X5), X6) = X3) & m2_nattra_1(X6, X0, X1, X4, X5)) & m2_cat_1(X5, X0, X1)) & m2_cat_1(X4, X0, X1)))), inference(usedef, [], [e216])).
fof(e216, plain, ! [X1, X0, X2] : (sP0(X1, X0, X2) <=> ! [X3] : (r2_hidden(X3, X2) <=> ? [X4] : (? [X5] : (? [X6] : (r2_nattra_1(X0, X1, X4, X5) & (k4_tarski(k4_tarski(X4, X5), X6) = X3) & m2_nattra_1(X6, X0, X1, X4, X5)) & m2_cat_1(X5, X0, X1)) & m2_cat_1(X4, X0, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f113, plain, ! [X0] : (! [X1] : (! [X2] : (((k11_nattra_1(X0, X1) = X2) <=> ! [X3] : (r2_hidden(X3, X2) <=> ? [X4] : (? [X5] : (? [X6] : (r2_nattra_1(X0, X1, X4, X5) & (k4_tarski(k4_tarski(X4, X5), X6) = X3) & m2_nattra_1(X6, X0, X1, X4, X5)) & m2_cat_1(X5, X0, X1)) & m2_cat_1(X4, X0, X1)))) | ~ m4_nattra_1(X2, X0, X1)) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(flattening, [], [f112])).
fof(f112, plain, ! [X0] : (! [X1] : (! [X2] : (((k11_nattra_1(X0, X1) = X2) <=> ! [X3] : (r2_hidden(X3, X2) <=> ? [X4] : (? [X5] : (? [X6] : (r2_nattra_1(X0, X1, X4, X5) & (k4_tarski(k4_tarski(X4, X5), X6) = X3) & m2_nattra_1(X6, X0, X1, X4, X5)) & m2_cat_1(X5, X0, X1)) & m2_cat_1(X4, X0, X1)))) | ~ m4_nattra_1(X2, X0, X1)) | (~ l1_cat_1(X1) | ~ v2_cat_1(X1))) | (~ l1_cat_1(X0) | ~ v2_cat_1(X0))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ! [X1] : ((l1_cat_1(X1) & v2_cat_1(X1)) => ! [X2] : (m4_nattra_1(X2, X0, X1) => ((k11_nattra_1(X0, X1) = X2) <=> ! [X3] : (r2_hidden(X3, X2) <=> ? [X4] : (? [X5] : (? [X6] : (r2_nattra_1(X0, X1, X4, X5) & (k4_tarski(k4_tarski(X4, X5), X6) = X3) & m2_nattra_1(X6, X0, X1, X4, X5)) & m2_cat_1(X5, X0, X1)) & m2_cat_1(X4, X0, X1))))))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', d16_nattra_1)).
fof(f38404, plain, (~ sP0(k1_yoneda_1(sK6), sK6, k11_nattra_1(sK6, k1_yoneda_1(sK6))) | (~ spl42_6 | ~ spl42_8 | spl42_1048)), inference(backward_demodulation, [], [f31986, f38402])).
fof(f38402, plain, ((u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6))) = k11_nattra_1(sK6, k1_yoneda_1(sK6))) | (~ spl42_6 | ~ spl42_8)), inference(subsumption_resolution, [], [f38364, f599])).
fof(f38364, plain, (~ v2_cat_1(k1_yoneda_1(sK6)) | (u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6))) = k11_nattra_1(sK6, k1_yoneda_1(sK6))) | ~ spl42_8), inference(resolution, [], [f4224, f609])).
fof(f4224, plain, ! [X30] : (~ l1_cat_1(X30) | ~ v2_cat_1(X30) | (k11_nattra_1(sK6, X30) = u2_cat_1(k12_nattra_1(sK6, X30)))), inference(subsumption_resolution, [], [f4204, f294])).
fof(f4204, plain, ! [X30] : (~ v2_cat_1(X30) | ~ v2_cat_1(sK6) | ~ l1_cat_1(X30) | (k11_nattra_1(sK6, X30) = u2_cat_1(k12_nattra_1(sK6, X30)))), inference(resolution, [], [f1641, f295])).
fof(f1641, plain, ! [X2, X3] : (~ l1_cat_1(X3) | ~ v2_cat_1(X2) | ~ v2_cat_1(X3) | ~ l1_cat_1(X2) | (k11_nattra_1(X3, X2) = u2_cat_1(k12_nattra_1(X3, X2)))), inference(resolution, [], [f842, f323])).
fof(f323, plain, ! [X2, X0, X1] : (~ sP4(X0, X1, X2) | (u2_cat_1(X2) = k11_nattra_1(X1, X0))), inference(cnf_transformation, [], [f247])).
fof(f247, plain, ! [X0, X1, X2] : ((sP4(X0, X1, X2) | ((~ (k10_cat_1(X2, sK15(X0, X1, X2)) = k4_tarski(k4_tarski(sK16(X0, X1, X2), sK16(X0, X1, X2)), k7_nattra_1(X1, X0, sK16(X0, X1, X2)))) & (sK15(X0, X1, X2) = sK16(X0, X1, X2)) & m2_cat_1(sK16(X0, X1, X2), X1, X0)) & m1_subset_1(sK15(X0, X1, X2), u1_cat_1(X2))) | ~ sP3(X0, X1, X2) | ~ sP2(X2) | ((~ (k3_cat_1(X2, sK17(X2)) = k2_mcart_1(k1_mcart_1(sK17(X2)))) | ~ (k2_cat_1(X2, sK17(X2)) = k1_mcart_1(k1_mcart_1(sK17(X2))))) & m1_subset_1(sK17(X2), u2_cat_1(X2))) | ~ (u2_cat_1(X2) = k11_nattra_1(X1, X0)) | ~ (u1_cat_1(X2) = k7_cat_2(X1, X0))) & ((! [X6] : (! [X7] : ((k10_cat_1(X2, X6) = k4_tarski(k4_tarski(X7, X7), k7_nattra_1(X1, X0, X7))) | ~ (X6 = X7) | ~ m2_cat_1(X7, X1, X0)) | ~ m1_subset_1(X6, u1_cat_1(X2))) & sP3(X0, X1, X2) & sP2(X2) & ! [X8] : (((k3_cat_1(X2, X8) = k2_mcart_1(k1_mcart_1(X8))) & (k2_cat_1(X2, X8) = k1_mcart_1(k1_mcart_1(X8)))) | ~ m1_subset_1(X8, u2_cat_1(X2))) & (u2_cat_1(X2) = k11_nattra_1(X1, X0)) & (u1_cat_1(X2) = k7_cat_2(X1, X0))) | ~ sP4(X0, X1, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16, sK17])], [f243, f246, f245, f244])).
fof(f244, plain, ! [X2, X1, X0] : (? [X3] : (? [X4] : (~ (k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X1, X0, X4))) & (X3 = X4) & m2_cat_1(X4, X1, X0)) & m1_subset_1(X3, u1_cat_1(X2))) => (? [X4] : (~ (k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X1, X0, X4)) = k10_cat_1(X2, sK15(X0, X1, X2))) & (sK15(X0, X1, X2) = X4) & m2_cat_1(X4, X1, X0)) & m1_subset_1(sK15(X0, X1, X2), u1_cat_1(X2)))), introduced(choice_axiom, [])).
fof(f245, plain, ! [X2, X1, X0] : (? [X4] : (~ (k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X1, X0, X4)) = k10_cat_1(X2, sK15(X0, X1, X2))) & (sK15(X0, X1, X2) = X4) & m2_cat_1(X4, X1, X0)) => (~ (k10_cat_1(X2, sK15(X0, X1, X2)) = k4_tarski(k4_tarski(sK16(X0, X1, X2), sK16(X0, X1, X2)), k7_nattra_1(X1, X0, sK16(X0, X1, X2)))) & (sK15(X0, X1, X2) = sK16(X0, X1, X2)) & m2_cat_1(sK16(X0, X1, X2), X1, X0))), introduced(choice_axiom, [])).
fof(f246, plain, ! [X2] : (? [X5] : ((~ (k3_cat_1(X2, X5) = k2_mcart_1(k1_mcart_1(X5))) | ~ (k2_cat_1(X2, X5) = k1_mcart_1(k1_mcart_1(X5)))) & m1_subset_1(X5, u2_cat_1(X2))) => ((~ (k3_cat_1(X2, sK17(X2)) = k2_mcart_1(k1_mcart_1(sK17(X2)))) | ~ (k2_cat_1(X2, sK17(X2)) = k1_mcart_1(k1_mcart_1(sK17(X2))))) & m1_subset_1(sK17(X2), u2_cat_1(X2)))), introduced(choice_axiom, [])).
fof(f243, plain, ! [X0, X1, X2] : ((sP4(X0, X1, X2) | ? [X3] : (? [X4] : (~ (k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X1, X0, X4))) & (X3 = X4) & m2_cat_1(X4, X1, X0)) & m1_subset_1(X3, u1_cat_1(X2))) | ~ sP3(X0, X1, X2) | ~ sP2(X2) | ? [X5] : ((~ (k3_cat_1(X2, X5) = k2_mcart_1(k1_mcart_1(X5))) | ~ (k2_cat_1(X2, X5) = k1_mcart_1(k1_mcart_1(X5)))) & m1_subset_1(X5, u2_cat_1(X2))) | ~ (u2_cat_1(X2) = k11_nattra_1(X1, X0)) | ~ (u1_cat_1(X2) = k7_cat_2(X1, X0))) & ((! [X6] : (! [X7] : ((k10_cat_1(X2, X6) = k4_tarski(k4_tarski(X7, X7), k7_nattra_1(X1, X0, X7))) | ~ (X6 = X7) | ~ m2_cat_1(X7, X1, X0)) | ~ m1_subset_1(X6, u1_cat_1(X2))) & sP3(X0, X1, X2) & sP2(X2) & ! [X8] : (((k3_cat_1(X2, X8) = k2_mcart_1(k1_mcart_1(X8))) & (k2_cat_1(X2, X8) = k1_mcart_1(k1_mcart_1(X8)))) | ~ m1_subset_1(X8, u2_cat_1(X2))) & (u2_cat_1(X2) = k11_nattra_1(X1, X0)) & (u1_cat_1(X2) = k7_cat_2(X1, X0))) | ~ sP4(X0, X1, X2))), inference(rectify, [], [f242])).
fof(f242, plain, ! [X1, X0, X2] : ((sP4(X1, X0, X2) | ? [X3] : (? [X4] : (~ (k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X0, X1, X4))) & (X3 = X4) & m2_cat_1(X4, X0, X1)) & m1_subset_1(X3, u1_cat_1(X2))) | ~ sP3(X1, X0, X2) | ~ sP2(X2) | ? [X14] : ((~ (k3_cat_1(X2, X14) = k2_mcart_1(k1_mcart_1(X14))) | ~ (k2_cat_1(X2, X14) = k1_mcart_1(k1_mcart_1(X14)))) & m1_subset_1(X14, u2_cat_1(X2))) | ~ (k11_nattra_1(X0, X1) = u2_cat_1(X2)) | ~ (u1_cat_1(X2) = k7_cat_2(X0, X1))) & ((! [X3] : (! [X4] : ((k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X0, X1, X4))) | ~ (X3 = X4) | ~ m2_cat_1(X4, X0, X1)) | ~ m1_subset_1(X3, u1_cat_1(X2))) & sP3(X1, X0, X2) & sP2(X2) & ! [X14] : (((k3_cat_1(X2, X14) = k2_mcart_1(k1_mcart_1(X14))) & (k2_cat_1(X2, X14) = k1_mcart_1(k1_mcart_1(X14)))) | ~ m1_subset_1(X14, u2_cat_1(X2))) & (k11_nattra_1(X0, X1) = u2_cat_1(X2)) & (u1_cat_1(X2) = k7_cat_2(X0, X1))) | ~ sP4(X1, X0, X2))), inference(flattening, [], [f241])).
fof(f241, plain, ! [X1, X0, X2] : ((sP4(X1, X0, X2) | (? [X3] : (? [X4] : (~ (k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X0, X1, X4))) & (X3 = X4) & m2_cat_1(X4, X0, X1)) & m1_subset_1(X3, u1_cat_1(X2))) | ~ sP3(X1, X0, X2) | ~ sP2(X2) | ? [X14] : ((~ (k3_cat_1(X2, X14) = k2_mcart_1(k1_mcart_1(X14))) | ~ (k2_cat_1(X2, X14) = k1_mcart_1(k1_mcart_1(X14)))) & m1_subset_1(X14, u2_cat_1(X2))) | ~ (k11_nattra_1(X0, X1) = u2_cat_1(X2)) | ~ (u1_cat_1(X2) = k7_cat_2(X0, X1)))) & ((! [X3] : (! [X4] : ((k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X0, X1, X4))) | ~ (X3 = X4) | ~ m2_cat_1(X4, X0, X1)) | ~ m1_subset_1(X3, u1_cat_1(X2))) & sP3(X1, X0, X2) & sP2(X2) & ! [X14] : (((k3_cat_1(X2, X14) = k2_mcart_1(k1_mcart_1(X14))) & (k2_cat_1(X2, X14) = k1_mcart_1(k1_mcart_1(X14)))) | ~ m1_subset_1(X14, u2_cat_1(X2))) & (k11_nattra_1(X0, X1) = u2_cat_1(X2)) & (u1_cat_1(X2) = k7_cat_2(X0, X1))) | ~ sP4(X1, X0, X2))), inference(nnf_transformation, [], [f221])).
fof(f221, plain, ! [X1, X0, X2] : (sP4(X1, X0, X2) <=> (! [X3] : (! [X4] : ((k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X0, X1, X4))) | ~ (X3 = X4) | ~ m2_cat_1(X4, X0, X1)) | ~ m1_subset_1(X3, u1_cat_1(X2))) & sP3(X1, X0, X2) & sP2(X2) & ! [X14] : (((k3_cat_1(X2, X14) = k2_mcart_1(k1_mcart_1(X14))) & (k2_cat_1(X2, X14) = k1_mcart_1(k1_mcart_1(X14)))) | ~ m1_subset_1(X14, u2_cat_1(X2))) & (k11_nattra_1(X0, X1) = u2_cat_1(X2)) & (u1_cat_1(X2) = k7_cat_2(X0, X1)))), inference(usedef, [], [e221])).
fof(e221, plain, ! [X1, X0, X2] : (sP4(X1, X0, X2) <=> (! [X3] : (! [X4] : ((k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X0, X1, X4))) | ~ (X3 = X4) | ~ m2_cat_1(X4, X0, X1)) | ~ m1_subset_1(X3, u1_cat_1(X2))) & sP3(X1, X0, X2) & sP2(X2) & ! [X14] : (((k3_cat_1(X2, X14) = k2_mcart_1(k1_mcart_1(X14))) & (k2_cat_1(X2, X14) = k1_mcart_1(k1_mcart_1(X14)))) | ~ m1_subset_1(X14, u2_cat_1(X2))) & (k11_nattra_1(X0, X1) = u2_cat_1(X2)) & (u1_cat_1(X2) = k7_cat_2(X0, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f842, plain, ! [X0, X1] : (sP4(X1, X0, k12_nattra_1(X0, X1)) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0) | ~ l1_cat_1(X1)), inference(subsumption_resolution, [], [f841, f365])).
fof(f365, plain, ! [X0, X1] : (v1_cat_1(k12_nattra_1(X0, X1)) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ! [X0, X1] : ((l1_cat_1(k12_nattra_1(X0, X1)) & v2_cat_1(k12_nattra_1(X0, X1)) & v1_cat_1(k12_nattra_1(X0, X1))) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(flattening, [], [f129])).
fof(f129, plain, ! [X0, X1] : ((l1_cat_1(k12_nattra_1(X0, X1)) & v2_cat_1(k12_nattra_1(X0, X1)) & v1_cat_1(k12_nattra_1(X0, X1))) | (~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ! [X0, X1] : ((l1_cat_1(X1) & v2_cat_1(X1) & l1_cat_1(X0) & v2_cat_1(X0)) => (l1_cat_1(k12_nattra_1(X0, X1)) & v2_cat_1(k12_nattra_1(X0, X1)) & v1_cat_1(k12_nattra_1(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', dt_k12_nattra_1)).
fof(f841, plain, ! [X0, X1] : (~ v1_cat_1(k12_nattra_1(X0, X1)) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0) | sP4(X1, X0, k12_nattra_1(X0, X1))), inference(subsumption_resolution, [], [f840, f366])).
fof(f366, plain, ! [X0, X1] : (v2_cat_1(k12_nattra_1(X0, X1)) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f130])).
fof(f840, plain, ! [X0, X1] : (~ v2_cat_1(k12_nattra_1(X0, X1)) | ~ v1_cat_1(k12_nattra_1(X0, X1)) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0) | sP4(X1, X0, k12_nattra_1(X0, X1))), inference(subsumption_resolution, [], [f838, f367])).
fof(f367, plain, ! [X0, X1] : (l1_cat_1(k12_nattra_1(X0, X1)) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f130])).
fof(f838, plain, ! [X0, X1] : (~ l1_cat_1(k12_nattra_1(X0, X1)) | ~ v2_cat_1(k12_nattra_1(X0, X1)) | ~ v1_cat_1(k12_nattra_1(X0, X1)) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0) | sP4(X1, X0, k12_nattra_1(X0, X1))), inference(resolution, [], [f354, f471])).
fof(f471, plain, ! [X2, X1] : (~ sP5(k12_nattra_1(X1, X2), X1, X2) | sP4(X2, X1, k12_nattra_1(X1, X2))), inference(equality_resolution, [], [f320])).
fof(f320, plain, ! [X2, X0, X1] : (sP4(X2, X1, X0) | ~ (k12_nattra_1(X1, X2) = X0) | ~ sP5(X0, X1, X2)), inference(cnf_transformation, [], [f240])).
fof(f240, plain, ! [X0, X1, X2] : ((((k12_nattra_1(X1, X2) = X0) | ~ sP4(X2, X1, X0)) & (sP4(X2, X1, X0) | ~ (k12_nattra_1(X1, X2) = X0))) | ~ sP5(X0, X1, X2)), inference(rectify, [], [f239])).
fof(f239, plain, ! [X2, X0, X1] : ((((k12_nattra_1(X0, X1) = X2) | ~ sP4(X1, X0, X2)) & (sP4(X1, X0, X2) | ~ (k12_nattra_1(X0, X1) = X2))) | ~ sP5(X2, X0, X1)), inference(nnf_transformation, [], [f222])).
fof(f222, plain, ! [X2, X0, X1] : (((k12_nattra_1(X0, X1) = X2) <=> sP4(X1, X0, X2)) | ~ sP5(X2, X0, X1)), inference(usedef, [], [e222])).
fof(e222, plain, ! [X2, X0, X1] : (sP5(X2, X0, X1) <=> ((k12_nattra_1(X0, X1) = X2) <=> sP4(X1, X0, X2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f354, plain, ! [X2, X0, X1] : (sP5(X2, X0, X1) | ~ l1_cat_1(X2) | ~ v2_cat_1(X2) | ~ v1_cat_1(X2) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f223])).
fof(f223, plain, ! [X0] : (! [X1] : (! [X2] : (sP5(X2, X0, X1) | ~ l1_cat_1(X2) | ~ v2_cat_1(X2) | ~ v1_cat_1(X2)) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(definition_folding, [], [f115, e222, e221, e220, e219])).
fof(f219, plain, ! [X2] : (sP2(X2) <=> ! [X12] : (! [X13] : (r2_hidden(k13_cat_2(X2, X2, X13, X12), k1_relat_1(u5_cat_1(X2))) | ~ (k3_cat_1(X2, X12) = k2_cat_1(X2, X13)) | ~ m1_subset_1(X13, u2_cat_1(X2))) | ~ m1_subset_1(X12, u2_cat_1(X2)))), inference(usedef, [], [e219])).
fof(e219, plain, ! [X2] : (sP2(X2) <=> ! [X12] : (! [X13] : (r2_hidden(k13_cat_2(X2, X2, X13, X12), k1_relat_1(u5_cat_1(X2))) | ~ (k3_cat_1(X2, X12) = k2_cat_1(X2, X13)) | ~ m1_subset_1(X13, u2_cat_1(X2))) | ~ m1_subset_1(X12, u2_cat_1(X2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f220, plain, ! [X1, X0, X2] : (sP3(X1, X0, X2) <=> ! [X5] : (! [X6] : (? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : ((k1_funct_1(u5_cat_1(X2), k13_cat_2(X2, X2, X6, X5)) = k4_tarski(k4_tarski(X7, X9), k8_nattra_1(X0, X1, X7, X8, X9, X10, X11))) & (k4_tarski(k4_tarski(X8, X9), X11) = X6) & (k4_tarski(k4_tarski(X7, X8), X10) = X5) & m2_nattra_1(X11, X0, X1, X8, X9)) & m2_nattra_1(X10, X0, X1, X7, X8)) & m2_cat_1(X9, X0, X1)) & m2_cat_1(X8, X0, X1)) & m2_cat_1(X7, X0, X1)) | ~ r2_hidden(k13_cat_2(X2, X2, X6, X5), k1_relat_1(u5_cat_1(X2))) | ~ m1_subset_1(X6, u2_cat_1(X2))) | ~ m1_subset_1(X5, u2_cat_1(X2)))), inference(usedef, [], [e220])).
fof(e220, plain, ! [X1, X0, X2] : (sP3(X1, X0, X2) <=> ! [X5] : (! [X6] : (? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : ((k1_funct_1(u5_cat_1(X2), k13_cat_2(X2, X2, X6, X5)) = k4_tarski(k4_tarski(X7, X9), k8_nattra_1(X0, X1, X7, X8, X9, X10, X11))) & (k4_tarski(k4_tarski(X8, X9), X11) = X6) & (k4_tarski(k4_tarski(X7, X8), X10) = X5) & m2_nattra_1(X11, X0, X1, X8, X9)) & m2_nattra_1(X10, X0, X1, X7, X8)) & m2_cat_1(X9, X0, X1)) & m2_cat_1(X8, X0, X1)) & m2_cat_1(X7, X0, X1)) | ~ r2_hidden(k13_cat_2(X2, X2, X6, X5), k1_relat_1(u5_cat_1(X2))) | ~ m1_subset_1(X6, u2_cat_1(X2))) | ~ m1_subset_1(X5, u2_cat_1(X2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f115, plain, ! [X0] : (! [X1] : (! [X2] : (((k12_nattra_1(X0, X1) = X2) <=> (! [X3] : (! [X4] : ((k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X0, X1, X4))) | ~ (X3 = X4) | ~ m2_cat_1(X4, X0, X1)) | ~ m1_subset_1(X3, u1_cat_1(X2))) & ! [X5] : (! [X6] : (? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : ((k1_funct_1(u5_cat_1(X2), k13_cat_2(X2, X2, X6, X5)) = k4_tarski(k4_tarski(X7, X9), k8_nattra_1(X0, X1, X7, X8, X9, X10, X11))) & (k4_tarski(k4_tarski(X8, X9), X11) = X6) & (k4_tarski(k4_tarski(X7, X8), X10) = X5) & m2_nattra_1(X11, X0, X1, X8, X9)) & m2_nattra_1(X10, X0, X1, X7, X8)) & m2_cat_1(X9, X0, X1)) & m2_cat_1(X8, X0, X1)) & m2_cat_1(X7, X0, X1)) | ~ r2_hidden(k13_cat_2(X2, X2, X6, X5), k1_relat_1(u5_cat_1(X2))) | ~ m1_subset_1(X6, u2_cat_1(X2))) | ~ m1_subset_1(X5, u2_cat_1(X2))) & ! [X12] : (! [X13] : (r2_hidden(k13_cat_2(X2, X2, X13, X12), k1_relat_1(u5_cat_1(X2))) | ~ (k3_cat_1(X2, X12) = k2_cat_1(X2, X13)) | ~ m1_subset_1(X13, u2_cat_1(X2))) | ~ m1_subset_1(X12, u2_cat_1(X2))) & ! [X14] : (((k3_cat_1(X2, X14) = k2_mcart_1(k1_mcart_1(X14))) & (k2_cat_1(X2, X14) = k1_mcart_1(k1_mcart_1(X14)))) | ~ m1_subset_1(X14, u2_cat_1(X2))) & (k11_nattra_1(X0, X1) = u2_cat_1(X2)) & (u1_cat_1(X2) = k7_cat_2(X0, X1)))) | ~ l1_cat_1(X2) | ~ v2_cat_1(X2) | ~ v1_cat_1(X2)) | ~ l1_cat_1(X1) | ~ v2_cat_1(X1)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(flattening, [], [f114])).
fof(f114, plain, ! [X0] : (! [X1] : (! [X2] : (((k12_nattra_1(X0, X1) = X2) <=> (! [X3] : (! [X4] : (((k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X0, X1, X4))) | ~ (X3 = X4)) | ~ m2_cat_1(X4, X0, X1)) | ~ m1_subset_1(X3, u1_cat_1(X2))) & ! [X5] : (! [X6] : ((? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : (((k1_funct_1(u5_cat_1(X2), k13_cat_2(X2, X2, X6, X5)) = k4_tarski(k4_tarski(X7, X9), k8_nattra_1(X0, X1, X7, X8, X9, X10, X11))) & (k4_tarski(k4_tarski(X8, X9), X11) = X6) & (k4_tarski(k4_tarski(X7, X8), X10) = X5)) & m2_nattra_1(X11, X0, X1, X8, X9)) & m2_nattra_1(X10, X0, X1, X7, X8)) & m2_cat_1(X9, X0, X1)) & m2_cat_1(X8, X0, X1)) & m2_cat_1(X7, X0, X1)) | ~ r2_hidden(k13_cat_2(X2, X2, X6, X5), k1_relat_1(u5_cat_1(X2)))) | ~ m1_subset_1(X6, u2_cat_1(X2))) | ~ m1_subset_1(X5, u2_cat_1(X2))) & ! [X12] : (! [X13] : ((r2_hidden(k13_cat_2(X2, X2, X13, X12), k1_relat_1(u5_cat_1(X2))) | ~ (k3_cat_1(X2, X12) = k2_cat_1(X2, X13))) | ~ m1_subset_1(X13, u2_cat_1(X2))) | ~ m1_subset_1(X12, u2_cat_1(X2))) & ! [X14] : (((k3_cat_1(X2, X14) = k2_mcart_1(k1_mcart_1(X14))) & (k2_cat_1(X2, X14) = k1_mcart_1(k1_mcart_1(X14)))) | ~ m1_subset_1(X14, u2_cat_1(X2))) & (k11_nattra_1(X0, X1) = u2_cat_1(X2)) & (u1_cat_1(X2) = k7_cat_2(X0, X1)))) | (~ l1_cat_1(X2) | ~ v2_cat_1(X2) | ~ v1_cat_1(X2))) | (~ l1_cat_1(X1) | ~ v2_cat_1(X1))) | (~ l1_cat_1(X0) | ~ v2_cat_1(X0))), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ! [X1] : ((l1_cat_1(X1) & v2_cat_1(X1)) => ! [X2] : ((l1_cat_1(X2) & v2_cat_1(X2) & v1_cat_1(X2)) => ((k12_nattra_1(X0, X1) = X2) <=> (! [X3] : (m1_subset_1(X3, u1_cat_1(X2)) => ! [X4] : (m2_cat_1(X4, X0, X1) => ((X3 = X4) => (k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X0, X1, X4)))))) & ! [X5] : (m1_subset_1(X5, u2_cat_1(X2)) => ! [X6] : (m1_subset_1(X6, u2_cat_1(X2)) => ~ (! [X7] : (m2_cat_1(X7, X0, X1) => ! [X8] : (m2_cat_1(X8, X0, X1) => ! [X9] : (m2_cat_1(X9, X0, X1) => ! [X10] : (m2_nattra_1(X10, X0, X1, X7, X8) => ! [X11] : (m2_nattra_1(X11, X0, X1, X8, X9) => ~ ((k1_funct_1(u5_cat_1(X2), k13_cat_2(X2, X2, X6, X5)) = k4_tarski(k4_tarski(X7, X9), k8_nattra_1(X0, X1, X7, X8, X9, X10, X11))) & (k4_tarski(k4_tarski(X8, X9), X11) = X6) & (k4_tarski(k4_tarski(X7, X8), X10) = X5))))))) & r2_hidden(k13_cat_2(X2, X2, X6, X5), k1_relat_1(u5_cat_1(X2)))))) & ! [X12] : (m1_subset_1(X12, u2_cat_1(X2)) => ! [X13] : (m1_subset_1(X13, u2_cat_1(X2)) => ((k3_cat_1(X2, X12) = k2_cat_1(X2, X13)) => r2_hidden(k13_cat_2(X2, X2, X13, X12), k1_relat_1(u5_cat_1(X2)))))) & ! [X14] : (m1_subset_1(X14, u2_cat_1(X2)) => ((k3_cat_1(X2, X14) = k2_mcart_1(k1_mcart_1(X14))) & (k2_cat_1(X2, X14) = k1_mcart_1(k1_mcart_1(X14))))) & (k11_nattra_1(X0, X1) = u2_cat_1(X2)) & (u1_cat_1(X2) = k7_cat_2(X0, X1))))))), inference(rectify, [], [f10])).
fof(f10, plain, ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ! [X1] : ((l1_cat_1(X1) & v2_cat_1(X1)) => ! [X2] : ((l1_cat_1(X2) & v2_cat_1(X2) & v1_cat_1(X2)) => ((k12_nattra_1(X0, X1) = X2) <=> (! [X3] : (m1_subset_1(X3, u1_cat_1(X2)) => ! [X4] : (m2_cat_1(X4, X0, X1) => ((X3 = X4) => (k10_cat_1(X2, X3) = k4_tarski(k4_tarski(X4, X4), k7_nattra_1(X0, X1, X4)))))) & ! [X3] : (m1_subset_1(X3, u2_cat_1(X2)) => ! [X4] : (m1_subset_1(X4, u2_cat_1(X2)) => ~ (! [X5] : (m2_cat_1(X5, X0, X1) => ! [X6] : (m2_cat_1(X6, X0, X1) => ! [X7] : (m2_cat_1(X7, X0, X1) => ! [X8] : (m2_nattra_1(X8, X0, X1, X5, X6) => ! [X9] : (m2_nattra_1(X9, X0, X1, X6, X7) => ~ ((k1_funct_1(u5_cat_1(X2), k13_cat_2(X2, X2, X4, X3)) = k4_tarski(k4_tarski(X5, X7), k8_nattra_1(X0, X1, X5, X6, X7, X8, X9))) & (k4_tarski(k4_tarski(X6, X7), X9) = X4) & (k4_tarski(k4_tarski(X5, X6), X8) = X3))))))) & r2_hidden(k13_cat_2(X2, X2, X4, X3), k1_relat_1(u5_cat_1(X2)))))) & ! [X3] : (m1_subset_1(X3, u2_cat_1(X2)) => ! [X4] : (m1_subset_1(X4, u2_cat_1(X2)) => ((k3_cat_1(X2, X3) = k2_cat_1(X2, X4)) => r2_hidden(k13_cat_2(X2, X2, X4, X3), k1_relat_1(u5_cat_1(X2)))))) & ! [X3] : (m1_subset_1(X3, u2_cat_1(X2)) => ((k3_cat_1(X2, X3) = k2_mcart_1(k1_mcart_1(X3))) & (k2_cat_1(X2, X3) = k1_mcart_1(k1_mcart_1(X3))))) & (k11_nattra_1(X0, X1) = u2_cat_1(X2)) & (u1_cat_1(X2) = k7_cat_2(X0, X1))))))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', d18_nattra_1)).
fof(f31986, plain, (~ sP0(k1_yoneda_1(sK6), sK6, u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))) | spl42_1048), inference(avatar_component_clause, [], [f31984])).
fof(f31984, plain, (spl42_1048 <=> sP0(k1_yoneda_1(sK6), sK6, u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6))))), introduced(avatar_definition, [new_symbols(naming, [spl42_1048])])).
fof(f36625, plain, (~ spl42_66 | spl42_1049), inference(avatar_contradiction_clause, [], [f36624])).
fof(f36624, plain, ($false | (~ spl42_66 | spl42_1049)), inference(subsumption_resolution, [], [f36623, f294])).
fof(f36623, plain, (~ v2_cat_1(sK6) | (~ spl42_66 | spl42_1049)), inference(subsumption_resolution, [], [f36622, f295])).
fof(f36622, plain, (~ l1_cat_1(sK6) | ~ v2_cat_1(sK6) | (~ spl42_66 | spl42_1049)), inference(subsumption_resolution, [], [f36621, f2888])).
fof(f2888, plain, (m1_subset_1(k3_cat_1(sK6, sK7), u1_cat_1(sK6)) | ~ spl42_66), inference(avatar_component_clause, [], [f2887])).
fof(f2887, plain, (spl42_66 <=> m1_subset_1(k3_cat_1(sK6, sK7), u1_cat_1(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl42_66])])).
fof(f36621, plain, (~ m1_subset_1(k3_cat_1(sK6, sK7), u1_cat_1(sK6)) | ~ l1_cat_1(sK6) | ~ v2_cat_1(sK6) | spl42_1049), inference(resolution, [], [f31990, f375])).
fof(f375, plain, ! [X0, X1] : (m2_cat_1(k2_yoneda_1(X0, X1), X0, k1_yoneda_1(X0)) | ~ m1_subset_1(X1, u1_cat_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ! [X0, X1] : (m2_cat_1(k2_yoneda_1(X0, X1), X0, k1_yoneda_1(X0)) | ~ m1_subset_1(X1, u1_cat_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(flattening, [], [f139])).
fof(f139, plain, ! [X0, X1] : (m2_cat_1(k2_yoneda_1(X0, X1), X0, k1_yoneda_1(X0)) | (~ m1_subset_1(X1, u1_cat_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0))), inference(ennf_transformation, [], [f34])).
fof(f34, plain, ! [X0, X1] : ((m1_subset_1(X1, u1_cat_1(X0)) & l1_cat_1(X0) & v2_cat_1(X0)) => m2_cat_1(k2_yoneda_1(X0, X1), X0, k1_yoneda_1(X0))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', dt_k2_yoneda_1)).
fof(f31990, plain, (~ m2_cat_1(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6)) | spl42_1049), inference(avatar_component_clause, [], [f31988])).
fof(f31988, plain, (spl42_1049 <=> m2_cat_1(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl42_1049])])).
fof(f36378, plain, (~ spl42_45 | spl42_1047), inference(avatar_contradiction_clause, [], [f36377])).
fof(f36377, plain, ($false | (~ spl42_45 | spl42_1047)), inference(subsumption_resolution, [], [f36376, f294])).
fof(f36376, plain, (~ v2_cat_1(sK6) | (~ spl42_45 | spl42_1047)), inference(subsumption_resolution, [], [f36375, f295])).
fof(f36375, plain, (~ l1_cat_1(sK6) | ~ v2_cat_1(sK6) | (~ spl42_45 | spl42_1047)), inference(subsumption_resolution, [], [f36374, f2126])).
fof(f2126, plain, (m1_subset_1(k2_cat_1(sK6, sK7), u1_cat_1(sK6)) | ~ spl42_45), inference(avatar_component_clause, [], [f2125])).
fof(f2125, plain, (spl42_45 <=> m1_subset_1(k2_cat_1(sK6, sK7), u1_cat_1(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl42_45])])).
fof(f36374, plain, (~ m1_subset_1(k2_cat_1(sK6, sK7), u1_cat_1(sK6)) | ~ l1_cat_1(sK6) | ~ v2_cat_1(sK6) | spl42_1047), inference(resolution, [], [f31982, f375])).
fof(f31982, plain, (~ m2_cat_1(k2_yoneda_1(sK6, k2_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6)) | spl42_1047), inference(avatar_component_clause, [], [f31980])).
fof(f31980, plain, (spl42_1047 <=> m2_cat_1(k2_yoneda_1(sK6, k2_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl42_1047])])).
fof(f31991, plain, (~ spl42_1047 | ~ spl42_1048 | ~ spl42_1049), inference(avatar_split_clause, [], [f31978, f31988, f31984, f31980])).
fof(f31978, plain, (~ m2_cat_1(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6)) | ~ sP0(k1_yoneda_1(sK6), sK6, u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))) | ~ m2_cat_1(k2_yoneda_1(sK6, k2_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6))), inference(subsumption_resolution, [], [f31977, f294])).
fof(f31977, plain, (~ m2_cat_1(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6)) | ~ sP0(k1_yoneda_1(sK6), sK6, u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))) | ~ v2_cat_1(sK6) | ~ m2_cat_1(k2_yoneda_1(sK6, k2_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6))), inference(subsumption_resolution, [], [f31976, f295])).
fof(f31976, plain, (~ m2_cat_1(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6)) | ~ sP0(k1_yoneda_1(sK6), sK6, u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))) | ~ l1_cat_1(sK6) | ~ v2_cat_1(sK6) | ~ m2_cat_1(k2_yoneda_1(sK6, k2_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6))), inference(subsumption_resolution, [], [f31880, f296])).
fof(f296, plain, m1_subset_1(sK7, u2_cat_1(sK6)), inference(cnf_transformation, [], [f226])).
fof(f31880, plain, (~ m2_cat_1(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6)) | ~ sP0(k1_yoneda_1(sK6), sK6, u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))) | ~ m1_subset_1(sK7, u2_cat_1(sK6)) | ~ l1_cat_1(sK6) | ~ v2_cat_1(sK6) | ~ m2_cat_1(k2_yoneda_1(sK6, k2_cat_1(sK6, sK7)), sK6, k1_yoneda_1(sK6))), inference(resolution, [], [f3078, f475])).
fof(f475, plain, ~ m1_subset_1(k2_tarski(k1_tarski(k2_tarski(k1_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7))), k2_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7))))), k2_tarski(k3_yoneda_1(sK6, sK7), k2_tarski(k1_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7))), k2_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7)))))), u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))), inference(forward_demodulation, [], [f474, f304])).
fof(f304, plain, ! [X0, X1] : (k2_tarski(X0, X1) = k2_tarski(X1, X0)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1] : (k2_tarski(X0, X1) = k2_tarski(X1, X0)), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', commutativity_k2_tarski)).
fof(f474, plain, ~ m1_subset_1(k2_tarski(k1_tarski(k2_tarski(k2_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7))), k1_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7))))), k2_tarski(k3_yoneda_1(sK6, sK7), k2_tarski(k2_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7))), k1_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)))))), u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))), inference(forward_demodulation, [], [f473, f304])).
fof(f473, plain, ~ m1_subset_1(k2_tarski(k1_tarski(k2_tarski(k2_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7))), k1_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7))))), k2_tarski(k2_tarski(k2_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7))), k1_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)))), k3_yoneda_1(sK6, sK7))), u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))), inference(backward_demodulation, [], [f456, f304])).
fof(f456, plain, ~ m1_subset_1(k2_tarski(k2_tarski(k2_tarski(k2_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7))), k1_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)))), k3_yoneda_1(sK6, sK7)), k1_tarski(k2_tarski(k2_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7))), k1_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)))))), u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))), inference(definition_unfolding, [], [f297, f357, f357])).
fof(f357, plain, ! [X0, X1] : (k4_tarski(X0, X1) = k2_tarski(k2_tarski(X0, X1), k1_tarski(X0))), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ! [X0, X1] : (k4_tarski(X0, X1) = k2_tarski(k2_tarski(X0, X1), k1_tarski(X0))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', d5_tarski)).
fof(f297, plain, ~ m1_subset_1(k4_tarski(k4_tarski(k2_yoneda_1(sK6, k3_cat_1(sK6, sK7)), k2_yoneda_1(sK6, k2_cat_1(sK6, sK7))), k3_yoneda_1(sK6, sK7)), u2_cat_1(k12_nattra_1(sK6, k1_yoneda_1(sK6)))), inference(cnf_transformation, [], [f226])).
fof(f3078, plain, ! [X6, X8, X7] : (m1_subset_1(k2_tarski(k1_tarski(k2_tarski(k1_tarski(k2_yoneda_1(X6, k3_cat_1(X6, X7))), k2_tarski(k2_yoneda_1(X6, k3_cat_1(X6, X7)), k2_yoneda_1(X6, k2_cat_1(X6, X7))))), k2_tarski(k3_yoneda_1(X6, X7), k2_tarski(k1_tarski(k2_yoneda_1(X6, k3_cat_1(X6, X7))), k2_tarski(k2_yoneda_1(X6, k3_cat_1(X6, X7)), k2_yoneda_1(X6, k2_cat_1(X6, X7)))))), X8) | ~ m2_cat_1(k2_yoneda_1(X6, k3_cat_1(X6, X7)), X6, k1_yoneda_1(X6)) | ~ sP0(k1_yoneda_1(X6), X6, X8) | ~ m1_subset_1(X7, u2_cat_1(X6)) | ~ l1_cat_1(X6) | ~ v2_cat_1(X6) | ~ m2_cat_1(k2_yoneda_1(X6, k2_cat_1(X6, X7)), X6, k1_yoneda_1(X6))), inference(resolution, [], [f1106, f447])).
fof(f447, plain, ! [X0, X1] : (~ r2_hidden(X0, X1) | m1_subset_1(X0, X1)), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ! [X0, X1] : (m1_subset_1(X0, X1) | ~ r2_hidden(X0, X1)), inference(ennf_transformation, [], [f89])).
fof(f89, plain, ! [X0, X1] : (r2_hidden(X0, X1) => m1_subset_1(X0, X1)), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', t1_subset)).
fof(f1106, plain, ! [X2, X0, X1] : (r2_hidden(k2_tarski(k1_tarski(k2_tarski(k1_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1))), k2_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))))), k2_tarski(k3_yoneda_1(X0, X1), k2_tarski(k1_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1))), k2_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1)))))), X2) | ~ m2_cat_1(k2_yoneda_1(X0, k2_cat_1(X0, X1)), X0, k1_yoneda_1(X0)) | ~ m2_cat_1(k2_yoneda_1(X0, k3_cat_1(X0, X1)), X0, k1_yoneda_1(X0)) | ~ sP0(k1_yoneda_1(X0), X0, X2) | ~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(forward_demodulation, [], [f1105, f304])).
fof(f1105, plain, ! [X2, X0, X1] : (r2_hidden(k2_tarski(k1_tarski(k2_tarski(k1_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1))), k2_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))))), k2_tarski(k2_tarski(k1_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1))), k2_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1)))), k3_yoneda_1(X0, X1))), X2) | ~ m2_cat_1(k2_yoneda_1(X0, k2_cat_1(X0, X1)), X0, k1_yoneda_1(X0)) | ~ m2_cat_1(k2_yoneda_1(X0, k3_cat_1(X0, X1)), X0, k1_yoneda_1(X0)) | ~ sP0(k1_yoneda_1(X0), X0, X2) | ~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(subsumption_resolution, [], [f1099, f450])).
fof(f450, plain, ! [X0, X1] : (r2_nattra_1(X0, k1_yoneda_1(X0), k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))) | ~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f209])).
fof(f209, plain, ! [X0] : (! [X1] : (r2_nattra_1(X0, k1_yoneda_1(X0), k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))) | ~ m1_subset_1(X1, u2_cat_1(X0))) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(flattening, [], [f208])).
fof(f208, plain, ! [X0] : (! [X1] : (r2_nattra_1(X0, k1_yoneda_1(X0), k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))) | ~ m1_subset_1(X1, u2_cat_1(X0))) | (~ l1_cat_1(X0) | ~ v2_cat_1(X0))), inference(ennf_transformation, [], [f92])).
fof(f92, plain, ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ! [X1] : (m1_subset_1(X1, u2_cat_1(X0)) => r2_nattra_1(X0, k1_yoneda_1(X0), k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', t3_yoneda_1)).
fof(f1099, plain, ! [X2, X0, X1] : (~ r2_nattra_1(X0, k1_yoneda_1(X0), k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))) | r2_hidden(k2_tarski(k1_tarski(k2_tarski(k1_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1))), k2_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))))), k2_tarski(k2_tarski(k1_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1))), k2_tarski(k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1)))), k3_yoneda_1(X0, X1))), X2) | ~ m2_cat_1(k2_yoneda_1(X0, k2_cat_1(X0, X1)), X0, k1_yoneda_1(X0)) | ~ m2_cat_1(k2_yoneda_1(X0, k3_cat_1(X0, X1)), X0, k1_yoneda_1(X0)) | ~ sP0(k1_yoneda_1(X0), X0, X2) | ~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(resolution, [], [f482, f377])).
fof(f377, plain, ! [X0, X1] : (m2_nattra_1(k3_yoneda_1(X0, X1), X0, k1_yoneda_1(X0), k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))) | ~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ! [X0, X1] : (m2_nattra_1(k3_yoneda_1(X0, X1), X0, k1_yoneda_1(X0), k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))) | ~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(flattening, [], [f143])).
fof(f143, plain, ! [X0, X1] : (m2_nattra_1(k3_yoneda_1(X0, X1), X0, k1_yoneda_1(X0), k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1))) | (~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ! [X0, X1] : ((m1_subset_1(X1, u2_cat_1(X0)) & l1_cat_1(X0) & v2_cat_1(X0)) => m2_nattra_1(k3_yoneda_1(X0, X1), X0, k1_yoneda_1(X0), k2_yoneda_1(X0, k3_cat_1(X0, X1)), k2_yoneda_1(X0, k2_cat_1(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', dt_k3_yoneda_1)).
fof(f482, plain, ! [X2, X0, X12, X1, X13, X11] : (~ m2_nattra_1(X13, X1, X0, X11, X12) | ~ r2_nattra_1(X1, X0, X11, X12) | r2_hidden(k2_tarski(k1_tarski(k2_tarski(k1_tarski(X11), k2_tarski(X11, X12))), k2_tarski(k2_tarski(k1_tarski(X11), k2_tarski(X11, X12)), X13)), X2) | ~ m2_cat_1(X12, X1, X0) | ~ m2_cat_1(X11, X1, X0) | ~ sP0(X0, X1, X2)), inference(forward_demodulation, [], [f481, f304])).
fof(f481, plain, ! [X2, X0, X12, X1, X13, X11] : (r2_hidden(k2_tarski(k1_tarski(k2_tarski(k2_tarski(X11, X12), k1_tarski(X11))), k2_tarski(k2_tarski(k2_tarski(X11, X12), k1_tarski(X11)), X13)), X2) | ~ r2_nattra_1(X1, X0, X11, X12) | ~ m2_nattra_1(X13, X1, X0, X11, X12) | ~ m2_cat_1(X12, X1, X0) | ~ m2_cat_1(X11, X1, X0) | ~ sP0(X0, X1, X2)), inference(forward_demodulation, [], [f470, f304])).
fof(f470, plain, ! [X2, X0, X12, X1, X13, X11] : (r2_hidden(k2_tarski(k2_tarski(k2_tarski(k2_tarski(X11, X12), k1_tarski(X11)), X13), k1_tarski(k2_tarski(k2_tarski(X11, X12), k1_tarski(X11)))), X2) | ~ r2_nattra_1(X1, X0, X11, X12) | ~ m2_nattra_1(X13, X1, X0, X11, X12) | ~ m2_cat_1(X12, X1, X0) | ~ m2_cat_1(X11, X1, X0) | ~ sP0(X0, X1, X2)), inference(equality_resolution, [], [f459])).
fof(f459, plain, ! [X2, X0, X12, X10, X1, X13, X11] : (r2_hidden(X10, X2) | ~ r2_nattra_1(X1, X0, X11, X12) | ~ (k2_tarski(k2_tarski(k2_tarski(k2_tarski(X11, X12), k1_tarski(X11)), X13), k1_tarski(k2_tarski(k2_tarski(X11, X12), k1_tarski(X11)))) = X10) | ~ m2_nattra_1(X13, X1, X0, X11, X12) | ~ m2_cat_1(X12, X1, X0) | ~ m2_cat_1(X11, X1, X0) | ~ sP0(X0, X1, X2)), inference(definition_unfolding, [], [f312, f357, f357])).
fof(f312, plain, ! [X2, X0, X12, X10, X1, X13, X11] : (r2_hidden(X10, X2) | ~ r2_nattra_1(X1, X0, X11, X12) | ~ (k4_tarski(k4_tarski(X11, X12), X13) = X10) | ~ m2_nattra_1(X13, X1, X0, X11, X12) | ~ m2_cat_1(X12, X1, X0) | ~ m2_cat_1(X11, X1, X0) | ~ sP0(X0, X1, X2)), inference(cnf_transformation, [], [f238])).
fof(f238, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | ((! [X4] : (! [X5] : (! [X6] : (~ r2_nattra_1(X1, X0, X4, X5) | ~ (k4_tarski(k4_tarski(X4, X5), X6) = sK8(X0, X1, X2)) | ~ m2_nattra_1(X6, X1, X0, X4, X5)) | ~ m2_cat_1(X5, X1, X0)) | ~ m2_cat_1(X4, X1, X0)) | ~ r2_hidden(sK8(X0, X1, X2), X2)) & ((((r2_nattra_1(X1, X0, sK9(X0, X1, X2), sK10(X0, X1, X2)) & (sK8(X0, X1, X2) = k4_tarski(k4_tarski(sK9(X0, X1, X2), sK10(X0, X1, X2)), sK11(X0, X1, X2))) & m2_nattra_1(sK11(X0, X1, X2), X1, X0, sK9(X0, X1, X2), sK10(X0, X1, X2))) & m2_cat_1(sK10(X0, X1, X2), X1, X0)) & m2_cat_1(sK9(X0, X1, X2), X1, X0)) | r2_hidden(sK8(X0, X1, X2), X2)))) & (! [X10] : ((r2_hidden(X10, X2) | ! [X11] : (! [X12] : (! [X13] : (~ r2_nattra_1(X1, X0, X11, X12) | ~ (k4_tarski(k4_tarski(X11, X12), X13) = X10) | ~ m2_nattra_1(X13, X1, X0, X11, X12)) | ~ m2_cat_1(X12, X1, X0)) | ~ m2_cat_1(X11, X1, X0))) & ((((r2_nattra_1(X1, X0, sK12(X0, X1, X10), sK13(X0, X1, X10)) & (k4_tarski(k4_tarski(sK12(X0, X1, X10), sK13(X0, X1, X10)), sK14(X0, X1, X10)) = X10) & m2_nattra_1(sK14(X0, X1, X10), X1, X0, sK12(X0, X1, X10), sK13(X0, X1, X10))) & m2_cat_1(sK13(X0, X1, X10), X1, X0)) & m2_cat_1(sK12(X0, X1, X10), X1, X0)) | ~ r2_hidden(X10, X2))) | ~ sP0(X0, X1, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8, sK9, sK10, sK11, sK12, sK13, sK14])], [f230, f237, f236, f235, f234, f233, f232, f231])).
fof(f231, plain, ! [X2, X1, X0] : (? [X3] : ((! [X4] : (! [X5] : (! [X6] : (~ r2_nattra_1(X1, X0, X4, X5) | ~ (k4_tarski(k4_tarski(X4, X5), X6) = X3) | ~ m2_nattra_1(X6, X1, X0, X4, X5)) | ~ m2_cat_1(X5, X1, X0)) | ~ m2_cat_1(X4, X1, X0)) | ~ r2_hidden(X3, X2)) & (? [X7] : (? [X8] : (? [X9] : (r2_nattra_1(X1, X0, X7, X8) & (k4_tarski(k4_tarski(X7, X8), X9) = X3) & m2_nattra_1(X9, X1, X0, X7, X8)) & m2_cat_1(X8, X1, X0)) & m2_cat_1(X7, X1, X0)) | r2_hidden(X3, X2))) => ((! [X4] : (! [X5] : (! [X6] : (~ r2_nattra_1(X1, X0, X4, X5) | ~ (k4_tarski(k4_tarski(X4, X5), X6) = sK8(X0, X1, X2)) | ~ m2_nattra_1(X6, X1, X0, X4, X5)) | ~ m2_cat_1(X5, X1, X0)) | ~ m2_cat_1(X4, X1, X0)) | ~ r2_hidden(sK8(X0, X1, X2), X2)) & (? [X7] : (? [X8] : (? [X9] : (r2_nattra_1(X1, X0, X7, X8) & (k4_tarski(k4_tarski(X7, X8), X9) = sK8(X0, X1, X2)) & m2_nattra_1(X9, X1, X0, X7, X8)) & m2_cat_1(X8, X1, X0)) & m2_cat_1(X7, X1, X0)) | r2_hidden(sK8(X0, X1, X2), X2)))), introduced(choice_axiom, [])).
fof(f232, plain, ! [X2, X1, X0] : (? [X7] : (? [X8] : (? [X9] : (r2_nattra_1(X1, X0, X7, X8) & (k4_tarski(k4_tarski(X7, X8), X9) = sK8(X0, X1, X2)) & m2_nattra_1(X9, X1, X0, X7, X8)) & m2_cat_1(X8, X1, X0)) & m2_cat_1(X7, X1, X0)) => (? [X8] : (? [X9] : (r2_nattra_1(X1, X0, sK9(X0, X1, X2), X8) & (sK8(X0, X1, X2) = k4_tarski(k4_tarski(sK9(X0, X1, X2), X8), X9)) & m2_nattra_1(X9, X1, X0, sK9(X0, X1, X2), X8)) & m2_cat_1(X8, X1, X0)) & m2_cat_1(sK9(X0, X1, X2), X1, X0))), introduced(choice_axiom, [])).
fof(f233, plain, ! [X2, X1, X0] : (? [X8] : (? [X9] : (r2_nattra_1(X1, X0, sK9(X0, X1, X2), X8) & (sK8(X0, X1, X2) = k4_tarski(k4_tarski(sK9(X0, X1, X2), X8), X9)) & m2_nattra_1(X9, X1, X0, sK9(X0, X1, X2), X8)) & m2_cat_1(X8, X1, X0)) => (? [X9] : (r2_nattra_1(X1, X0, sK9(X0, X1, X2), sK10(X0, X1, X2)) & (sK8(X0, X1, X2) = k4_tarski(k4_tarski(sK9(X0, X1, X2), sK10(X0, X1, X2)), X9)) & m2_nattra_1(X9, X1, X0, sK9(X0, X1, X2), sK10(X0, X1, X2))) & m2_cat_1(sK10(X0, X1, X2), X1, X0))), introduced(choice_axiom, [])).
fof(f234, plain, ! [X2, X1, X0] : (? [X9] : (r2_nattra_1(X1, X0, sK9(X0, X1, X2), sK10(X0, X1, X2)) & (sK8(X0, X1, X2) = k4_tarski(k4_tarski(sK9(X0, X1, X2), sK10(X0, X1, X2)), X9)) & m2_nattra_1(X9, X1, X0, sK9(X0, X1, X2), sK10(X0, X1, X2))) => (r2_nattra_1(X1, X0, sK9(X0, X1, X2), sK10(X0, X1, X2)) & (sK8(X0, X1, X2) = k4_tarski(k4_tarski(sK9(X0, X1, X2), sK10(X0, X1, X2)), sK11(X0, X1, X2))) & m2_nattra_1(sK11(X0, X1, X2), X1, X0, sK9(X0, X1, X2), sK10(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f235, plain, ! [X10, X1, X0] : (? [X14] : (? [X15] : (? [X16] : (r2_nattra_1(X1, X0, X14, X15) & (k4_tarski(k4_tarski(X14, X15), X16) = X10) & m2_nattra_1(X16, X1, X0, X14, X15)) & m2_cat_1(X15, X1, X0)) & m2_cat_1(X14, X1, X0)) => (? [X15] : (? [X16] : (r2_nattra_1(X1, X0, sK12(X0, X1, X10), X15) & (k4_tarski(k4_tarski(sK12(X0, X1, X10), X15), X16) = X10) & m2_nattra_1(X16, X1, X0, sK12(X0, X1, X10), X15)) & m2_cat_1(X15, X1, X0)) & m2_cat_1(sK12(X0, X1, X10), X1, X0))), introduced(choice_axiom, [])).
fof(f236, plain, ! [X10, X1, X0] : (? [X15] : (? [X16] : (r2_nattra_1(X1, X0, sK12(X0, X1, X10), X15) & (k4_tarski(k4_tarski(sK12(X0, X1, X10), X15), X16) = X10) & m2_nattra_1(X16, X1, X0, sK12(X0, X1, X10), X15)) & m2_cat_1(X15, X1, X0)) => (? [X16] : (r2_nattra_1(X1, X0, sK12(X0, X1, X10), sK13(X0, X1, X10)) & (k4_tarski(k4_tarski(sK12(X0, X1, X10), sK13(X0, X1, X10)), X16) = X10) & m2_nattra_1(X16, X1, X0, sK12(X0, X1, X10), sK13(X0, X1, X10))) & m2_cat_1(sK13(X0, X1, X10), X1, X0))), introduced(choice_axiom, [])).
fof(f237, plain, ! [X10, X1, X0] : (? [X16] : (r2_nattra_1(X1, X0, sK12(X0, X1, X10), sK13(X0, X1, X10)) & (k4_tarski(k4_tarski(sK12(X0, X1, X10), sK13(X0, X1, X10)), X16) = X10) & m2_nattra_1(X16, X1, X0, sK12(X0, X1, X10), sK13(X0, X1, X10))) => (r2_nattra_1(X1, X0, sK12(X0, X1, X10), sK13(X0, X1, X10)) & (k4_tarski(k4_tarski(sK12(X0, X1, X10), sK13(X0, X1, X10)), sK14(X0, X1, X10)) = X10) & m2_nattra_1(sK14(X0, X1, X10), X1, X0, sK12(X0, X1, X10), sK13(X0, X1, X10)))), introduced(choice_axiom, [])).
fof(f230, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | ? [X3] : ((! [X4] : (! [X5] : (! [X6] : (~ r2_nattra_1(X1, X0, X4, X5) | ~ (k4_tarski(k4_tarski(X4, X5), X6) = X3) | ~ m2_nattra_1(X6, X1, X0, X4, X5)) | ~ m2_cat_1(X5, X1, X0)) | ~ m2_cat_1(X4, X1, X0)) | ~ r2_hidden(X3, X2)) & (? [X7] : (? [X8] : (? [X9] : (r2_nattra_1(X1, X0, X7, X8) & (k4_tarski(k4_tarski(X7, X8), X9) = X3) & m2_nattra_1(X9, X1, X0, X7, X8)) & m2_cat_1(X8, X1, X0)) & m2_cat_1(X7, X1, X0)) | r2_hidden(X3, X2)))) & (! [X10] : ((r2_hidden(X10, X2) | ! [X11] : (! [X12] : (! [X13] : (~ r2_nattra_1(X1, X0, X11, X12) | ~ (k4_tarski(k4_tarski(X11, X12), X13) = X10) | ~ m2_nattra_1(X13, X1, X0, X11, X12)) | ~ m2_cat_1(X12, X1, X0)) | ~ m2_cat_1(X11, X1, X0))) & (? [X14] : (? [X15] : (? [X16] : (r2_nattra_1(X1, X0, X14, X15) & (k4_tarski(k4_tarski(X14, X15), X16) = X10) & m2_nattra_1(X16, X1, X0, X14, X15)) & m2_cat_1(X15, X1, X0)) & m2_cat_1(X14, X1, X0)) | ~ r2_hidden(X10, X2))) | ~ sP0(X0, X1, X2))), inference(rectify, [], [f229])).
fof(f229, plain, ! [X1, X0, X2] : ((sP0(X1, X0, X2) | ? [X3] : ((! [X4] : (! [X5] : (! [X6] : (~ r2_nattra_1(X0, X1, X4, X5) | ~ (k4_tarski(k4_tarski(X4, X5), X6) = X3) | ~ m2_nattra_1(X6, X0, X1, X4, X5)) | ~ m2_cat_1(X5, X0, X1)) | ~ m2_cat_1(X4, X0, X1)) | ~ r2_hidden(X3, X2)) & (? [X4] : (? [X5] : (? [X6] : (r2_nattra_1(X0, X1, X4, X5) & (k4_tarski(k4_tarski(X4, X5), X6) = X3) & m2_nattra_1(X6, X0, X1, X4, X5)) & m2_cat_1(X5, X0, X1)) & m2_cat_1(X4, X0, X1)) | r2_hidden(X3, X2)))) & (! [X3] : ((r2_hidden(X3, X2) | ! [X4] : (! [X5] : (! [X6] : (~ r2_nattra_1(X0, X1, X4, X5) | ~ (k4_tarski(k4_tarski(X4, X5), X6) = X3) | ~ m2_nattra_1(X6, X0, X1, X4, X5)) | ~ m2_cat_1(X5, X0, X1)) | ~ m2_cat_1(X4, X0, X1))) & (? [X4] : (? [X5] : (? [X6] : (r2_nattra_1(X0, X1, X4, X5) & (k4_tarski(k4_tarski(X4, X5), X6) = X3) & m2_nattra_1(X6, X0, X1, X4, X5)) & m2_cat_1(X5, X0, X1)) & m2_cat_1(X4, X0, X1)) | ~ r2_hidden(X3, X2))) | ~ sP0(X1, X0, X2))), inference(nnf_transformation, [], [f216])).
fof(f2916, plain, spl42_66, inference(avatar_contradiction_clause, [], [f2915])).
fof(f2915, plain, ($false | spl42_66), inference(subsumption_resolution, [], [f2914, f295])).
fof(f2914, plain, (~ l1_cat_1(sK6) | spl42_66), inference(subsumption_resolution, [], [f2913, f296])).
fof(f2913, plain, (~ m1_subset_1(sK7, u2_cat_1(sK6)) | ~ l1_cat_1(sK6) | spl42_66), inference(resolution, [], [f2889, f376])).
fof(f376, plain, ! [X0, X1] : (m1_subset_1(k3_cat_1(X0, X1), u1_cat_1(X0)) | ~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0)), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ! [X0, X1] : (m1_subset_1(k3_cat_1(X0, X1), u1_cat_1(X0)) | ~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0)), inference(flattening, [], [f141])).
fof(f141, plain, ! [X0, X1] : (m1_subset_1(k3_cat_1(X0, X1), u1_cat_1(X0)) | (~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0))), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ! [X0, X1] : ((m1_subset_1(X1, u2_cat_1(X0)) & l1_cat_1(X0)) => m1_subset_1(k3_cat_1(X0, X1), u1_cat_1(X0))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', dt_k3_cat_1)).
fof(f2889, plain, (~ m1_subset_1(k3_cat_1(sK6, sK7), u1_cat_1(sK6)) | spl42_66), inference(avatar_component_clause, [], [f2887])).
fof(f2150, plain, spl42_45, inference(avatar_contradiction_clause, [], [f2149])).
fof(f2149, plain, ($false | spl42_45), inference(subsumption_resolution, [], [f2148, f295])).
fof(f2148, plain, (~ l1_cat_1(sK6) | spl42_45), inference(subsumption_resolution, [], [f2147, f296])).
fof(f2147, plain, (~ m1_subset_1(sK7, u2_cat_1(sK6)) | ~ l1_cat_1(sK6) | spl42_45), inference(resolution, [], [f2127, f374])).
fof(f374, plain, ! [X0, X1] : (m1_subset_1(k2_cat_1(X0, X1), u1_cat_1(X0)) | ~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0)), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ! [X0, X1] : (m1_subset_1(k2_cat_1(X0, X1), u1_cat_1(X0)) | ~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0)), inference(flattening, [], [f137])).
fof(f137, plain, ! [X0, X1] : (m1_subset_1(k2_cat_1(X0, X1), u1_cat_1(X0)) | (~ m1_subset_1(X1, u2_cat_1(X0)) | ~ l1_cat_1(X0))), inference(ennf_transformation, [], [f30])).
fof(f30, plain, ! [X0, X1] : ((m1_subset_1(X1, u2_cat_1(X0)) & l1_cat_1(X0)) => m1_subset_1(k2_cat_1(X0, X1), u1_cat_1(X0))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', dt_k2_cat_1)).
fof(f2127, plain, (~ m1_subset_1(k2_cat_1(sK6, sK7), u1_cat_1(sK6)) | spl42_45), inference(avatar_component_clause, [], [f2125])).
fof(f620, plain, ~ spl42_5, inference(avatar_contradiction_clause, [], [f619])).
fof(f619, plain, ($false | ~ spl42_5), inference(subsumption_resolution, [], [f618, f294])).
fof(f618, plain, (~ v2_cat_1(sK6) | ~ spl42_5), inference(subsumption_resolution, [], [f614, f295])).
fof(f614, plain, (~ l1_cat_1(sK6) | ~ v2_cat_1(sK6) | ~ spl42_5), inference(resolution, [], [f595, f423])).
fof(f423, plain, ! [X0] : (~ v1_xboole_0(k17_ens_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f191])).
fof(f191, plain, ! [X0] : (~ v1_xboole_0(k17_ens_1(X0)) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(flattening, [], [f190])).
fof(f190, plain, ! [X0] : (~ v1_xboole_0(k17_ens_1(X0)) | (~ l1_cat_1(X0) | ~ v2_cat_1(X0))), inference(ennf_transformation, [], [f75])).
fof(f75, plain, ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ~ v1_xboole_0(k17_ens_1(X0))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', fc5_ens_1)).
fof(f595, plain, (v1_xboole_0(k17_ens_1(sK6)) | ~ spl42_5), inference(avatar_component_clause, [], [f593])).
fof(f593, plain, (spl42_5 <=> v1_xboole_0(k17_ens_1(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl42_5])])).
fof(f610, plain, (spl42_5 | spl42_8), inference(avatar_split_clause, [], [f591, f607, f593])).
fof(f591, plain, (l1_cat_1(k1_yoneda_1(sK6)) | v1_xboole_0(k17_ens_1(sK6))), inference(superposition, [], [f364, f576])).
fof(f576, plain, (k1_yoneda_1(sK6) = k12_ens_1(k17_ens_1(sK6))), inference(subsumption_resolution, [], [f572, f294])).
fof(f572, plain, ((k1_yoneda_1(sK6) = k12_ens_1(k17_ens_1(sK6))) | ~ v2_cat_1(sK6)), inference(resolution, [], [f355, f295])).
fof(f355, plain, ! [X0] : (~ l1_cat_1(X0) | (k1_yoneda_1(X0) = k12_ens_1(k17_ens_1(X0))) | ~ v2_cat_1(X0)), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ! [X0] : ((k1_yoneda_1(X0) = k12_ens_1(k17_ens_1(X0))) | ~ l1_cat_1(X0) | ~ v2_cat_1(X0)), inference(flattening, [], [f116])).
fof(f116, plain, ! [X0] : ((k1_yoneda_1(X0) = k12_ens_1(k17_ens_1(X0))) | (~ l1_cat_1(X0) | ~ v2_cat_1(X0))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => (k1_yoneda_1(X0) = k12_ens_1(k17_ens_1(X0)))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', d1_yoneda_1)).
fof(f364, plain, ! [X0] : (l1_cat_1(k12_ens_1(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ! [X0] : (l1_cat_1(k12_ens_1(X0)) | v1_xboole_0(X0)), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0] : (~ v1_xboole_0(X0) => l1_cat_1(k12_ens_1(X0))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', dt_k12_ens_1)).
fof(f600, plain, (spl42_5 | spl42_6), inference(avatar_split_clause, [], [f589, f597, f593])).
fof(f589, plain, (v2_cat_1(k1_yoneda_1(sK6)) | v1_xboole_0(k17_ens_1(sK6))), inference(superposition, [], [f422, f576])).
fof(f422, plain, ! [X0] : (v2_cat_1(k12_ens_1(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f189])).
fof(f189, plain, ! [X0] : ((v2_cat_1(k12_ens_1(X0)) & v1_cat_1(k12_ens_1(X0))) | v1_xboole_0(X0)), inference(ennf_transformation, [], [f74])).
fof(f74, plain, ! [X0] : (~ v1_xboole_0(X0) => (v2_cat_1(k12_ens_1(X0)) & v1_cat_1(k12_ens_1(X0)))), file('/home/ubuntu/library/tptp/Problems/CAT/CAT034+1.p', fc4_ens_1)).