fof(f684, plain, $false, inference(avatar_sat_refutation, [], [f354, f372, f425, f449, f465, f621, f637, f639, f683])).
fof(f683, plain, (~ spl51_2 | ~ spl51_16), inference(avatar_contradiction_clause, [], [f682])).
fof(f682, plain, ($false | (~ spl51_2 | ~ spl51_16)), inference(subsumption_resolution, [], [f681, f236])).
fof(f236, plain, aElement0(xb), inference(cnf_transformation, [], [f17])).
fof(f17, plain, (aElement0(xc) & aElement0(xb) & aElement0(xa)), file('/home/ubuntu/library/tptp/Problems/COM/COM022+4.p', m__731)).
fof(f681, plain, (~ aElement0(xb) | (~ spl51_2 | ~ spl51_16)), inference(subsumption_resolution, [], [f680, f340])).
fof(f340, plain, ~ sP18(xb), inference(equality_resolution, [], [f260])).
fof(f260, plain, ! [X0] : (~ (xb = X0) | ~ sP18(X0)), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ! [X0] : ((~ sdtmndtasgtdt0(xb, xR, X0) & ~ sdtmndtplgtdt0(xb, xR, X0) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, X0) | ~ aReductOfIn0(X1, xb, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(X0, xb, xR) & ~ (xb = X0)) | ~ sP18(X0)), inference(rectify, [], [f126])).
fof(f126, plain, ! [X16] : ((~ sdtmndtasgtdt0(xb, xR, X16) & ~ sdtmndtplgtdt0(xb, xR, X16) & ! [X18] : (~ sdtmndtplgtdt0(X18, xR, X16) | ~ aReductOfIn0(X18, xb, xR) | ~ aElement0(X18)) & ~ aReductOfIn0(X16, xb, xR) & ~ (xb = X16)) | ~ sP18(X16)), inference(nnf_transformation, [], [f76])).
fof(f76, plain, ! [X16] : ((~ sdtmndtasgtdt0(xb, xR, X16) & ~ sdtmndtplgtdt0(xb, xR, X16) & ! [X18] : (~ sdtmndtplgtdt0(X18, xR, X16) | ~ aReductOfIn0(X18, xb, xR) | ~ aElement0(X18)) & ~ aReductOfIn0(X16, xb, xR) & ~ (xb = X16)) | ~ sP18(X16)), inference(usedef, [], [e76])).
fof(e76, plain, ! [X16] : (sP18(X16) <=> (~ sdtmndtasgtdt0(xb, xR, X16) & ~ sdtmndtplgtdt0(xb, xR, X16) & ! [X18] : (~ sdtmndtplgtdt0(X18, xR, X16) | ~ aReductOfIn0(X18, xb, xR) | ~ aElement0(X18)) & ~ aReductOfIn0(X16, xb, xR) & ~ (xb = X16))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f680, plain, (sP18(xb) | ~ aElement0(xb) | (~ spl51_2 | ~ spl51_16)), inference(resolution, [], [f648, f352])).
fof(f352, plain, (sdtmndtplgtdt0(xa, xR, xb) | ~ spl51_2), inference(avatar_component_clause, [], [f351])).
fof(f351, plain, (spl51_2 <=> sdtmndtplgtdt0(xa, xR, xb)), introduced(avatar_definition, [new_symbols(naming, [spl51_2])])).
fof(f648, plain, (! [X0] : (~ sdtmndtplgtdt0(xa, xR, X0) | sP18(X0) | ~ aElement0(X0)) | ~ spl51_16), inference(backward_demodulation, [], [f335, f420])).
fof(f420, plain, ((xa = xc) | ~ spl51_16), inference(avatar_component_clause, [], [f418])).
fof(f418, plain, (spl51_16 <=> (xa = xc)), introduced(avatar_definition, [new_symbols(naming, [spl51_16])])).
fof(f335, plain, ! [X0] : (~ sdtmndtplgtdt0(xc, xR, X0) | sP18(X0) | ~ aElement0(X0)), inference(cnf_transformation, [], [f171])).
fof(f171, plain, (! [X0] : ((~ sdtmndtasgtdt0(xc, xR, X0) & ~ sdtmndtplgtdt0(xc, xR, X0) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, X0) | ~ aReductOfIn0(X1, xc, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(X0, xc, xR) & ~ (xc = X0)) | sP18(X0) | ~ aElement0(X0)) & sdtmndtasgtdt0(xa, xR, xc) & ((sdtmndtplgtdt0(xa, xR, xc) & ((sdtmndtplgtdt0(sK49, xR, xc) & aReductOfIn0(sK49, xa, xR) & aElement0(sK49)) | aReductOfIn0(xc, xa, xR))) | (xa = xc)) & sdtmndtasgtdt0(xa, xR, xb) & ((sdtmndtplgtdt0(xa, xR, xb) & ((sdtmndtplgtdt0(sK50, xR, xb) & aReductOfIn0(sK50, xa, xR) & aElement0(sK50)) | aReductOfIn0(xb, xa, xR))) | (xa = xb)) & (sP16 | (~ sdtmndtplgtdt0(xa, xR, xc) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, xc) | ~ aReductOfIn0(X4, xa, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(xc, xa, xR)) | sP17)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK49, sK50])], [f168, f170, f169])).
fof(f169, plain, (? [X2] : (sdtmndtplgtdt0(X2, xR, xc) & aReductOfIn0(X2, xa, xR) & aElement0(X2)) => (sdtmndtplgtdt0(sK49, xR, xc) & aReductOfIn0(sK49, xa, xR) & aElement0(sK49))), introduced(choice_axiom, [])).
fof(f170, plain, (? [X3] : (sdtmndtplgtdt0(X3, xR, xb) & aReductOfIn0(X3, xa, xR) & aElement0(X3)) => (sdtmndtplgtdt0(sK50, xR, xb) & aReductOfIn0(sK50, xa, xR) & aElement0(sK50))), introduced(choice_axiom, [])).
fof(f168, plain, (! [X0] : ((~ sdtmndtasgtdt0(xc, xR, X0) & ~ sdtmndtplgtdt0(xc, xR, X0) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, X0) | ~ aReductOfIn0(X1, xc, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(X0, xc, xR) & ~ (xc = X0)) | sP18(X0) | ~ aElement0(X0)) & sdtmndtasgtdt0(xa, xR, xc) & ((sdtmndtplgtdt0(xa, xR, xc) & (? [X2] : (sdtmndtplgtdt0(X2, xR, xc) & aReductOfIn0(X2, xa, xR) & aElement0(X2)) | aReductOfIn0(xc, xa, xR))) | (xa = xc)) & sdtmndtasgtdt0(xa, xR, xb) & ((sdtmndtplgtdt0(xa, xR, xb) & (? [X3] : (sdtmndtplgtdt0(X3, xR, xb) & aReductOfIn0(X3, xa, xR) & aElement0(X3)) | aReductOfIn0(xb, xa, xR))) | (xa = xb)) & (sP16 | (~ sdtmndtplgtdt0(xa, xR, xc) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, xc) | ~ aReductOfIn0(X4, xa, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(xc, xa, xR)) | sP17)), inference(rectify, [], [f77])).
fof(f77, plain, (! [X16] : ((~ sdtmndtasgtdt0(xc, xR, X16) & ~ sdtmndtplgtdt0(xc, xR, X16) & ! [X17] : (~ sdtmndtplgtdt0(X17, xR, X16) | ~ aReductOfIn0(X17, xc, xR) | ~ aElement0(X17)) & ~ aReductOfIn0(X16, xc, xR) & ~ (xc = X16)) | sP18(X16) | ~ aElement0(X16)) & sdtmndtasgtdt0(xa, xR, xc) & ((sdtmndtplgtdt0(xa, xR, xc) & (? [X14] : (sdtmndtplgtdt0(X14, xR, xc) & aReductOfIn0(X14, xa, xR) & aElement0(X14)) | aReductOfIn0(xc, xa, xR))) | (xa = xc)) & sdtmndtasgtdt0(xa, xR, xb) & ((sdtmndtplgtdt0(xa, xR, xb) & (? [X15] : (sdtmndtplgtdt0(X15, xR, xb) & aReductOfIn0(X15, xa, xR) & aElement0(X15)) | aReductOfIn0(xb, xa, xR))) | (xa = xb)) & (sP16 | (~ sdtmndtplgtdt0(xa, xR, xc) & ! [X0] : (~ sdtmndtplgtdt0(X0, xR, xc) | ~ aReductOfIn0(X0, xa, xR) | ~ aElement0(X0)) & ~ aReductOfIn0(xc, xa, xR)) | sP17)), inference(definition_folding, [], [f53, e76, e75, e74, e73, e72, e71, e70, e69, e68, e67, e66])).
fof(f66, plain, ! [X5, X4] : ((sdtmndtplgtdt0(X4, xR, X5) & (? [X9] : (sdtmndtplgtdt0(X9, xR, X5) & aReductOfIn0(X9, X4, xR) & aElement0(X9)) | aReductOfIn0(X5, X4, xR))) | (X4 = X5) | ~ sP8(X5, X4)), inference(usedef, [], [e66])).
fof(e66, plain, ! [X5, X4] : (sP8(X5, X4) <=> ((sdtmndtplgtdt0(X4, xR, X5) & (? [X9] : (sdtmndtplgtdt0(X9, xR, X5) & aReductOfIn0(X9, X4, xR) & aElement0(X9)) | aReductOfIn0(X5, X4, xR))) | (X4 = X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f67, plain, ! [X5] : ((sdtmndtplgtdt0(xb, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, xb, xR) & aElement0(X7)) | aReductOfIn0(X5, xb, xR))) | (xb = X5) | ~ sP9(X5)), inference(usedef, [], [e67])).
fof(e67, plain, ! [X5] : (sP9(X5) <=> ((sdtmndtplgtdt0(xb, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, xb, xR) & aElement0(X7)) | aReductOfIn0(X5, xb, xR))) | (xb = X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f68, plain, ! [X5] : ((sdtmndtplgtdt0(xc, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, xc, xR) & aElement0(X6)) | aReductOfIn0(X5, xc, xR))) | (xc = X5) | ~ sP10(X5)), inference(usedef, [], [e68])).
fof(e68, plain, ! [X5] : (sP10(X5) <=> ((sdtmndtplgtdt0(xc, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, xc, xR) & aElement0(X6)) | aReductOfIn0(X5, xc, xR))) | (xc = X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f69, plain, ! [X4] : (? [X5] : (sdtmndtasgtdt0(xc, xR, X5) & sP10(X5) & sdtmndtasgtdt0(xb, xR, X5) & sP9(X5) & aNormalFormOfIn0(X5, X4, xR) & ! [X8] : ~ aReductOfIn0(X8, X5, xR) & sdtmndtasgtdt0(X4, xR, X5) & sP8(X5, X4) & aElement0(X5)) | ~ sP11(X4)), inference(usedef, [], [e69])).
fof(e69, plain, ! [X4] : (sP11(X4) <=> ? [X5] : (sdtmndtasgtdt0(xc, xR, X5) & sP10(X5) & sdtmndtasgtdt0(xb, xR, X5) & sP9(X5) & aNormalFormOfIn0(X5, X4, xR) & ! [X8] : ~ aReductOfIn0(X8, X5, xR) & sdtmndtasgtdt0(X4, xR, X5) & sP8(X5, X4) & aElement0(X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f70, plain, ! [X4, X2] : ((sdtmndtplgtdt0(X2, xR, X4) & (? [X11] : (sdtmndtplgtdt0(X11, xR, X4) & aReductOfIn0(X11, X2, xR) & aElement0(X11)) | aReductOfIn0(X4, X2, xR))) | (X2 = X4) | ~ sP12(X4, X2)), inference(usedef, [], [e70])).
fof(e70, plain, ! [X4, X2] : (sP12(X4, X2) <=> ((sdtmndtplgtdt0(X2, xR, X4) & (? [X11] : (sdtmndtplgtdt0(X11, xR, X4) & aReductOfIn0(X11, X2, xR) & aElement0(X11)) | aReductOfIn0(X4, X2, xR))) | (X2 = X4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f71, plain, ! [X4, X3] : ((sdtmndtplgtdt0(X3, xR, X4) & (? [X10] : (sdtmndtplgtdt0(X10, xR, X4) & aReductOfIn0(X10, X3, xR) & aElement0(X10)) | aReductOfIn0(X4, X3, xR))) | (X3 = X4) | ~ sP13(X4, X3)), inference(usedef, [], [e71])).
fof(e71, plain, ! [X4, X3] : (sP13(X4, X3) <=> ((sdtmndtplgtdt0(X3, xR, X4) & (? [X10] : (sdtmndtplgtdt0(X10, xR, X4) & aReductOfIn0(X10, X3, xR) & aElement0(X10)) | aReductOfIn0(X4, X3, xR))) | (X3 = X4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f72, plain, ! [X3, X2] : (? [X4] : (sP11(X4) & sdtmndtasgtdt0(X3, xR, X4) & sP13(X4, X3) & sdtmndtasgtdt0(X2, xR, X4) & sP12(X4, X2) & aElement0(X4)) | ~ sP14(X3, X2)), inference(usedef, [], [e72])).
fof(e72, plain, ! [X3, X2] : (sP14(X3, X2) <=> ? [X4] : (sP11(X4) & sdtmndtasgtdt0(X3, xR, X4) & sP13(X4, X3) & sdtmndtasgtdt0(X2, xR, X4) & sP12(X4, X2) & aElement0(X4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f73, plain, ! [X2] : (? [X3] : (sP14(X3, X2) & sdtmndtasgtdt0(X3, xR, xc) & ((sdtmndtplgtdt0(X3, xR, xc) & (? [X12] : (sdtmndtplgtdt0(X12, xR, xc) & aReductOfIn0(X12, X3, xR) & aElement0(X12)) | aReductOfIn0(xc, X3, xR))) | (xc = X3)) & aReductOfIn0(X3, xa, xR) & aElement0(X3)) | ~ sP15(X2)), inference(usedef, [], [e73])).
fof(e73, plain, ! [X2] : (sP15(X2) <=> ? [X3] : (sP14(X3, X2) & sdtmndtasgtdt0(X3, xR, xc) & ((sdtmndtplgtdt0(X3, xR, xc) & (? [X12] : (sdtmndtplgtdt0(X12, xR, xc) & aReductOfIn0(X12, X3, xR) & aElement0(X12)) | aReductOfIn0(xc, X3, xR))) | (xc = X3)) & aReductOfIn0(X3, xa, xR) & aElement0(X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f74, plain, (? [X2] : (sP15(X2) & sdtmndtasgtdt0(X2, xR, xb) & ((sdtmndtplgtdt0(X2, xR, xb) & (? [X13] : (sdtmndtplgtdt0(X13, xR, xb) & aReductOfIn0(X13, X2, xR) & aElement0(X13)) | aReductOfIn0(xb, X2, xR))) | (xb = X2)) & aReductOfIn0(X2, xa, xR) & aElement0(X2)) | ~ sP16), inference(usedef, [], [e74])).
fof(e74, plain, (sP16 <=> ? [X2] : (sP15(X2) & sdtmndtasgtdt0(X2, xR, xb) & ((sdtmndtplgtdt0(X2, xR, xb) & (? [X13] : (sdtmndtplgtdt0(X13, xR, xb) & aReductOfIn0(X13, X2, xR) & aElement0(X13)) | aReductOfIn0(xb, X2, xR))) | (xb = X2)) & aReductOfIn0(X2, xa, xR) & aElement0(X2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f75, plain, ((~ sdtmndtplgtdt0(xa, xR, xb) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, xb) | ~ aReductOfIn0(X1, xa, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(xb, xa, xR)) | ~ sP17), inference(usedef, [], [e75])).
fof(e75, plain, (sP17 <=> (~ sdtmndtplgtdt0(xa, xR, xb) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, xb) | ~ aReductOfIn0(X1, xa, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(xb, xa, xR))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f53, plain, (! [X16] : ((~ sdtmndtasgtdt0(xc, xR, X16) & ~ sdtmndtplgtdt0(xc, xR, X16) & ! [X17] : (~ sdtmndtplgtdt0(X17, xR, X16) | ~ aReductOfIn0(X17, xc, xR) | ~ aElement0(X17)) & ~ aReductOfIn0(X16, xc, xR) & ~ (xc = X16)) | (~ sdtmndtasgtdt0(xb, xR, X16) & ~ sdtmndtplgtdt0(xb, xR, X16) & ! [X18] : (~ sdtmndtplgtdt0(X18, xR, X16) | ~ aReductOfIn0(X18, xb, xR) | ~ aElement0(X18)) & ~ aReductOfIn0(X16, xb, xR) & ~ (xb = X16)) | ~ aElement0(X16)) & sdtmndtasgtdt0(xa, xR, xc) & ((sdtmndtplgtdt0(xa, xR, xc) & (? [X14] : (sdtmndtplgtdt0(X14, xR, xc) & aReductOfIn0(X14, xa, xR) & aElement0(X14)) | aReductOfIn0(xc, xa, xR))) | (xa = xc)) & sdtmndtasgtdt0(xa, xR, xb) & ((sdtmndtplgtdt0(xa, xR, xb) & (? [X15] : (sdtmndtplgtdt0(X15, xR, xb) & aReductOfIn0(X15, xa, xR) & aElement0(X15)) | aReductOfIn0(xb, xa, xR))) | (xa = xb)) & (? [X2] : (? [X3] : (? [X4] : (? [X5] : (sdtmndtasgtdt0(xc, xR, X5) & ((sdtmndtplgtdt0(xc, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, xc, xR) & aElement0(X6)) | aReductOfIn0(X5, xc, xR))) | (xc = X5)) & sdtmndtasgtdt0(xb, xR, X5) & ((sdtmndtplgtdt0(xb, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, xb, xR) & aElement0(X7)) | aReductOfIn0(X5, xb, xR))) | (xb = X5)) & aNormalFormOfIn0(X5, X4, xR) & ! [X8] : ~ aReductOfIn0(X8, X5, xR) & sdtmndtasgtdt0(X4, xR, X5) & ((sdtmndtplgtdt0(X4, xR, X5) & (? [X9] : (sdtmndtplgtdt0(X9, xR, X5) & aReductOfIn0(X9, X4, xR) & aElement0(X9)) | aReductOfIn0(X5, X4, xR))) | (X4 = X5)) & aElement0(X5)) & sdtmndtasgtdt0(X3, xR, X4) & ((sdtmndtplgtdt0(X3, xR, X4) & (? [X10] : (sdtmndtplgtdt0(X10, xR, X4) & aReductOfIn0(X10, X3, xR) & aElement0(X10)) | aReductOfIn0(X4, X3, xR))) | (X3 = X4)) & sdtmndtasgtdt0(X2, xR, X4) & ((sdtmndtplgtdt0(X2, xR, X4) & (? [X11] : (sdtmndtplgtdt0(X11, xR, X4) & aReductOfIn0(X11, X2, xR) & aElement0(X11)) | aReductOfIn0(X4, X2, xR))) | (X2 = X4)) & aElement0(X4)) & sdtmndtasgtdt0(X3, xR, xc) & ((sdtmndtplgtdt0(X3, xR, xc) & (? [X12] : (sdtmndtplgtdt0(X12, xR, xc) & aReductOfIn0(X12, X3, xR) & aElement0(X12)) | aReductOfIn0(xc, X3, xR))) | (xc = X3)) & aReductOfIn0(X3, xa, xR) & aElement0(X3)) & sdtmndtasgtdt0(X2, xR, xb) & ((sdtmndtplgtdt0(X2, xR, xb) & (? [X13] : (sdtmndtplgtdt0(X13, xR, xb) & aReductOfIn0(X13, X2, xR) & aElement0(X13)) | aReductOfIn0(xb, X2, xR))) | (xb = X2)) & aReductOfIn0(X2, xa, xR) & aElement0(X2)) | (~ sdtmndtplgtdt0(xa, xR, xc) & ! [X0] : (~ sdtmndtplgtdt0(X0, xR, xc) | ~ aReductOfIn0(X0, xa, xR) | ~ aElement0(X0)) & ~ aReductOfIn0(xc, xa, xR)) | (~ sdtmndtplgtdt0(xa, xR, xb) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, xb) | ~ aReductOfIn0(X1, xa, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(xb, xa, xR)))), inference(flattening, [], [f52])).
fof(f52, plain, ((! [X16] : ((~ sdtmndtasgtdt0(xc, xR, X16) & ~ sdtmndtplgtdt0(xc, xR, X16) & ! [X17] : (~ sdtmndtplgtdt0(X17, xR, X16) | ~ aReductOfIn0(X17, xc, xR) | ~ aElement0(X17)) & ~ aReductOfIn0(X16, xc, xR) & ~ (xc = X16)) | (~ sdtmndtasgtdt0(xb, xR, X16) & ~ sdtmndtplgtdt0(xb, xR, X16) & ! [X18] : (~ sdtmndtplgtdt0(X18, xR, X16) | ~ aReductOfIn0(X18, xb, xR) | ~ aElement0(X18)) & ~ aReductOfIn0(X16, xb, xR) & ~ (xb = X16)) | ~ aElement0(X16)) & (sdtmndtasgtdt0(xa, xR, xc) & ((sdtmndtplgtdt0(xa, xR, xc) & (? [X14] : (sdtmndtplgtdt0(X14, xR, xc) & aReductOfIn0(X14, xa, xR) & aElement0(X14)) | aReductOfIn0(xc, xa, xR))) | (xa = xc)) & sdtmndtasgtdt0(xa, xR, xb) & ((sdtmndtplgtdt0(xa, xR, xb) & (? [X15] : (sdtmndtplgtdt0(X15, xR, xb) & aReductOfIn0(X15, xa, xR) & aElement0(X15)) | aReductOfIn0(xb, xa, xR))) | (xa = xb)))) & (? [X2] : (? [X3] : (? [X4] : (? [X5] : (sdtmndtasgtdt0(xc, xR, X5) & ((sdtmndtplgtdt0(xc, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, xc, xR) & aElement0(X6)) | aReductOfIn0(X5, xc, xR))) | (xc = X5)) & sdtmndtasgtdt0(xb, xR, X5) & ((sdtmndtplgtdt0(xb, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, xb, xR) & aElement0(X7)) | aReductOfIn0(X5, xb, xR))) | (xb = X5)) & aNormalFormOfIn0(X5, X4, xR) & ! [X8] : ~ aReductOfIn0(X8, X5, xR) & sdtmndtasgtdt0(X4, xR, X5) & ((sdtmndtplgtdt0(X4, xR, X5) & (? [X9] : (sdtmndtplgtdt0(X9, xR, X5) & aReductOfIn0(X9, X4, xR) & aElement0(X9)) | aReductOfIn0(X5, X4, xR))) | (X4 = X5)) & aElement0(X5)) & sdtmndtasgtdt0(X3, xR, X4) & ((sdtmndtplgtdt0(X3, xR, X4) & (? [X10] : (sdtmndtplgtdt0(X10, xR, X4) & aReductOfIn0(X10, X3, xR) & aElement0(X10)) | aReductOfIn0(X4, X3, xR))) | (X3 = X4)) & sdtmndtasgtdt0(X2, xR, X4) & ((sdtmndtplgtdt0(X2, xR, X4) & (? [X11] : (sdtmndtplgtdt0(X11, xR, X4) & aReductOfIn0(X11, X2, xR) & aElement0(X11)) | aReductOfIn0(X4, X2, xR))) | (X2 = X4)) & aElement0(X4)) & sdtmndtasgtdt0(X3, xR, xc) & ((sdtmndtplgtdt0(X3, xR, xc) & (? [X12] : (sdtmndtplgtdt0(X12, xR, xc) & aReductOfIn0(X12, X3, xR) & aElement0(X12)) | aReductOfIn0(xc, X3, xR))) | (xc = X3)) & aReductOfIn0(X3, xa, xR) & aElement0(X3)) & sdtmndtasgtdt0(X2, xR, xb) & ((sdtmndtplgtdt0(X2, xR, xb) & (? [X13] : (sdtmndtplgtdt0(X13, xR, xb) & aReductOfIn0(X13, X2, xR) & aElement0(X13)) | aReductOfIn0(xb, X2, xR))) | (xb = X2)) & aReductOfIn0(X2, xa, xR) & aElement0(X2)) | ((~ sdtmndtplgtdt0(xa, xR, xc) & ! [X0] : (~ sdtmndtplgtdt0(X0, xR, xc) | ~ aReductOfIn0(X0, xa, xR) | ~ aElement0(X0)) & ~ aReductOfIn0(xc, xa, xR)) | (~ sdtmndtplgtdt0(xa, xR, xb) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, xb) | ~ aReductOfIn0(X1, xa, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(xb, xa, xR))))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ ((((sdtmndtplgtdt0(xa, xR, xc) | ? [X0] : (sdtmndtplgtdt0(X0, xR, xc) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | aReductOfIn0(xc, xa, xR)) & (sdtmndtplgtdt0(xa, xR, xb) | ? [X1] : (sdtmndtplgtdt0(X1, xR, xb) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) | aReductOfIn0(xb, xa, xR))) => ? [X2] : (? [X3] : (? [X4] : (? [X5] : (sdtmndtasgtdt0(xc, xR, X5) & ((sdtmndtplgtdt0(xc, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, xc, xR) & aElement0(X6)) | aReductOfIn0(X5, xc, xR))) | (xc = X5)) & sdtmndtasgtdt0(xb, xR, X5) & ((sdtmndtplgtdt0(xb, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, xb, xR) & aElement0(X7)) | aReductOfIn0(X5, xb, xR))) | (xb = X5)) & aNormalFormOfIn0(X5, X4, xR) & ~ ? [X8] : aReductOfIn0(X8, X5, xR) & sdtmndtasgtdt0(X4, xR, X5) & ((sdtmndtplgtdt0(X4, xR, X5) & (? [X9] : (sdtmndtplgtdt0(X9, xR, X5) & aReductOfIn0(X9, X4, xR) & aElement0(X9)) | aReductOfIn0(X5, X4, xR))) | (X4 = X5)) & aElement0(X5)) & sdtmndtasgtdt0(X3, xR, X4) & ((sdtmndtplgtdt0(X3, xR, X4) & (? [X10] : (sdtmndtplgtdt0(X10, xR, X4) & aReductOfIn0(X10, X3, xR) & aElement0(X10)) | aReductOfIn0(X4, X3, xR))) | (X3 = X4)) & sdtmndtasgtdt0(X2, xR, X4) & ((sdtmndtplgtdt0(X2, xR, X4) & (? [X11] : (sdtmndtplgtdt0(X11, xR, X4) & aReductOfIn0(X11, X2, xR) & aElement0(X11)) | aReductOfIn0(X4, X2, xR))) | (X2 = X4)) & aElement0(X4)) & sdtmndtasgtdt0(X3, xR, xc) & ((sdtmndtplgtdt0(X3, xR, xc) & (? [X12] : (sdtmndtplgtdt0(X12, xR, xc) & aReductOfIn0(X12, X3, xR) & aElement0(X12)) | aReductOfIn0(xc, X3, xR))) | (xc = X3)) & aReductOfIn0(X3, xa, xR) & aElement0(X3)) & sdtmndtasgtdt0(X2, xR, xb) & ((sdtmndtplgtdt0(X2, xR, xb) & (? [X13] : (sdtmndtplgtdt0(X13, xR, xb) & aReductOfIn0(X13, X2, xR) & aElement0(X13)) | aReductOfIn0(xb, X2, xR))) | (xb = X2)) & aReductOfIn0(X2, xa, xR) & aElement0(X2))) => ((sdtmndtasgtdt0(xa, xR, xc) & ((sdtmndtplgtdt0(xa, xR, xc) & (? [X14] : (sdtmndtplgtdt0(X14, xR, xc) & aReductOfIn0(X14, xa, xR) & aElement0(X14)) | aReductOfIn0(xc, xa, xR))) | (xa = xc)) & sdtmndtasgtdt0(xa, xR, xb) & ((sdtmndtplgtdt0(xa, xR, xb) & (? [X15] : (sdtmndtplgtdt0(X15, xR, xb) & aReductOfIn0(X15, xa, xR) & aElement0(X15)) | aReductOfIn0(xb, xa, xR))) | (xa = xb))) => ? [X16] : ((sdtmndtasgtdt0(xc, xR, X16) | sdtmndtplgtdt0(xc, xR, X16) | ? [X17] : (sdtmndtplgtdt0(X17, xR, X16) & aReductOfIn0(X17, xc, xR) & aElement0(X17)) | aReductOfIn0(X16, xc, xR) | (xc = X16)) & (sdtmndtasgtdt0(xb, xR, X16) | sdtmndtplgtdt0(xb, xR, X16) | ? [X18] : (sdtmndtplgtdt0(X18, xR, X16) & aReductOfIn0(X18, xb, xR) & aElement0(X18)) | aReductOfIn0(X16, xb, xR) | (xb = X16)) & aElement0(X16)))), inference(rectify, [], [f20])).
fof(f20, plain, ~ ((((sdtmndtplgtdt0(xa, xR, xc) | ? [X0] : (sdtmndtplgtdt0(X0, xR, xc) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | aReductOfIn0(xc, xa, xR)) & (sdtmndtplgtdt0(xa, xR, xb) | ? [X0] : (sdtmndtplgtdt0(X0, xR, xb) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | aReductOfIn0(xb, xa, xR))) => ? [X0] : (? [X1] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & ((sdtmndtplgtdt0(xc, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, xc, xR) & aElement0(X4)) | aReductOfIn0(X3, xc, xR))) | (xc = X3)) & sdtmndtasgtdt0(xb, xR, X3) & ((sdtmndtplgtdt0(xb, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, xb, xR) & aElement0(X4)) | aReductOfIn0(X3, xb, xR))) | (xb = X3)) & aNormalFormOfIn0(X3, X2, xR) & ~ ? [X4] : aReductOfIn0(X4, X3, xR) & sdtmndtasgtdt0(X2, xR, X3) & ((sdtmndtplgtdt0(X2, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X2, xR) & aElement0(X4)) | aReductOfIn0(X3, X2, xR))) | (X2 = X3)) & aElement0(X3)) & sdtmndtasgtdt0(X1, xR, X2) & ((sdtmndtplgtdt0(X1, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X1, xR) & aElement0(X3)) | aReductOfIn0(X2, X1, xR))) | (X1 = X2)) & sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & aElement0(X2)) & sdtmndtasgtdt0(X1, xR, xc) & ((sdtmndtplgtdt0(X1, xR, xc) & (? [X2] : (sdtmndtplgtdt0(X2, xR, xc) & aReductOfIn0(X2, X1, xR) & aElement0(X2)) | aReductOfIn0(xc, X1, xR))) | (xc = X1)) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) & sdtmndtasgtdt0(X0, xR, xb) & ((sdtmndtplgtdt0(X0, xR, xb) & (? [X1] : (sdtmndtplgtdt0(X1, xR, xb) & aReductOfIn0(X1, X0, xR) & aElement0(X1)) | aReductOfIn0(xb, X0, xR))) | (xb = X0)) & aReductOfIn0(X0, xa, xR) & aElement0(X0))) => ((sdtmndtasgtdt0(xa, xR, xc) & ((sdtmndtplgtdt0(xa, xR, xc) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xc) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | aReductOfIn0(xc, xa, xR))) | (xa = xc)) & sdtmndtasgtdt0(xa, xR, xb) & ((sdtmndtplgtdt0(xa, xR, xb) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xb) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | aReductOfIn0(xb, xa, xR))) | (xa = xb))) => ? [X0] : ((sdtmndtasgtdt0(xc, xR, X0) | sdtmndtplgtdt0(xc, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xc, xR) & aElement0(X1)) | aReductOfIn0(X0, xc, xR) | (xc = X0)) & (sdtmndtasgtdt0(xb, xR, X0) | sdtmndtplgtdt0(xb, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xb, xR) & aElement0(X1)) | aReductOfIn0(X0, xb, xR) | (xb = X0)) & aElement0(X0)))), inference(negated_conjecture, [], [f19])).
fof(f19, plain, ~ ((((sdtmndtplgtdt0(xa, xR, xc) | ? [X0] : (sdtmndtplgtdt0(X0, xR, xc) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | aReductOfIn0(xc, xa, xR)) & (sdtmndtplgtdt0(xa, xR, xb) | ? [X0] : (sdtmndtplgtdt0(X0, xR, xb) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | aReductOfIn0(xb, xa, xR))) => ? [X0] : (? [X1] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & ((sdtmndtplgtdt0(xc, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, xc, xR) & aElement0(X4)) | aReductOfIn0(X3, xc, xR))) | (xc = X3)) & sdtmndtasgtdt0(xb, xR, X3) & ((sdtmndtplgtdt0(xb, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, xb, xR) & aElement0(X4)) | aReductOfIn0(X3, xb, xR))) | (xb = X3)) & aNormalFormOfIn0(X3, X2, xR) & ~ ? [X4] : aReductOfIn0(X4, X3, xR) & sdtmndtasgtdt0(X2, xR, X3) & ((sdtmndtplgtdt0(X2, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X2, xR) & aElement0(X4)) | aReductOfIn0(X3, X2, xR))) | (X2 = X3)) & aElement0(X3)) & sdtmndtasgtdt0(X1, xR, X2) & ((sdtmndtplgtdt0(X1, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X1, xR) & aElement0(X3)) | aReductOfIn0(X2, X1, xR))) | (X1 = X2)) & sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & aElement0(X2)) & sdtmndtasgtdt0(X1, xR, xc) & ((sdtmndtplgtdt0(X1, xR, xc) & (? [X2] : (sdtmndtplgtdt0(X2, xR, xc) & aReductOfIn0(X2, X1, xR) & aElement0(X2)) | aReductOfIn0(xc, X1, xR))) | (xc = X1)) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) & sdtmndtasgtdt0(X0, xR, xb) & ((sdtmndtplgtdt0(X0, xR, xb) & (? [X1] : (sdtmndtplgtdt0(X1, xR, xb) & aReductOfIn0(X1, X0, xR) & aElement0(X1)) | aReductOfIn0(xb, X0, xR))) | (xb = X0)) & aReductOfIn0(X0, xa, xR) & aElement0(X0))) => ((sdtmndtasgtdt0(xa, xR, xc) & ((sdtmndtplgtdt0(xa, xR, xc) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xc) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | aReductOfIn0(xc, xa, xR))) | (xa = xc)) & sdtmndtasgtdt0(xa, xR, xb) & ((sdtmndtplgtdt0(xa, xR, xb) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xb) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | aReductOfIn0(xb, xa, xR))) | (xa = xb))) => ? [X0] : ((sdtmndtasgtdt0(xc, xR, X0) | sdtmndtplgtdt0(xc, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xc, xR) & aElement0(X1)) | aReductOfIn0(X0, xc, xR) | (xc = X0)) & (sdtmndtasgtdt0(xb, xR, X0) | sdtmndtplgtdt0(xb, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xb, xR) & aElement0(X1)) | aReductOfIn0(X0, xb, xR) | (xb = X0)) & aElement0(X0)))), file('/home/ubuntu/library/tptp/Problems/COM/COM022+4.p', m__)).
fof(f639, plain, ~ spl51_6, inference(avatar_contradiction_clause, [], [f638])).
fof(f638, plain, ($false | ~ spl51_6), inference(subsumption_resolution, [], [f371, f561])).
fof(f561, plain, ! [X0] : ~ sP15(X0), inference(resolution, [], [f558, f283])).
fof(f283, plain, ! [X0] : (sP14(sK40(X0), X0) | ~ sP15(X0)), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ! [X0] : ((sP14(sK40(X0), X0) & sdtmndtasgtdt0(sK40(X0), xR, xc) & ((sdtmndtplgtdt0(sK40(X0), xR, xc) & ((sdtmndtplgtdt0(sK41(X0), xR, xc) & aReductOfIn0(sK41(X0), sK40(X0), xR) & aElement0(sK41(X0))) | aReductOfIn0(xc, sK40(X0), xR))) | (xc = sK40(X0))) & aReductOfIn0(sK40(X0), xa, xR) & aElement0(sK40(X0))) | ~ sP15(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK40, sK41])], [f136, f138, f137])).
fof(f137, plain, ! [X0] : (? [X1] : (sP14(X1, X0) & sdtmndtasgtdt0(X1, xR, xc) & ((sdtmndtplgtdt0(X1, xR, xc) & (? [X2] : (sdtmndtplgtdt0(X2, xR, xc) & aReductOfIn0(X2, X1, xR) & aElement0(X2)) | aReductOfIn0(xc, X1, xR))) | (xc = X1)) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) => (sP14(sK40(X0), X0) & sdtmndtasgtdt0(sK40(X0), xR, xc) & ((sdtmndtplgtdt0(sK40(X0), xR, xc) & (? [X2] : (sdtmndtplgtdt0(X2, xR, xc) & aReductOfIn0(X2, sK40(X0), xR) & aElement0(X2)) | aReductOfIn0(xc, sK40(X0), xR))) | (xc = sK40(X0))) & aReductOfIn0(sK40(X0), xa, xR) & aElement0(sK40(X0)))), introduced(choice_axiom, [])).
fof(f138, plain, ! [X0] : (? [X2] : (sdtmndtplgtdt0(X2, xR, xc) & aReductOfIn0(X2, sK40(X0), xR) & aElement0(X2)) => (sdtmndtplgtdt0(sK41(X0), xR, xc) & aReductOfIn0(sK41(X0), sK40(X0), xR) & aElement0(sK41(X0)))), introduced(choice_axiom, [])).
fof(f136, plain, ! [X0] : (? [X1] : (sP14(X1, X0) & sdtmndtasgtdt0(X1, xR, xc) & ((sdtmndtplgtdt0(X1, xR, xc) & (? [X2] : (sdtmndtplgtdt0(X2, xR, xc) & aReductOfIn0(X2, X1, xR) & aElement0(X2)) | aReductOfIn0(xc, X1, xR))) | (xc = X1)) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) | ~ sP15(X0)), inference(rectify, [], [f135])).
fof(f135, plain, ! [X2] : (? [X3] : (sP14(X3, X2) & sdtmndtasgtdt0(X3, xR, xc) & ((sdtmndtplgtdt0(X3, xR, xc) & (? [X12] : (sdtmndtplgtdt0(X12, xR, xc) & aReductOfIn0(X12, X3, xR) & aElement0(X12)) | aReductOfIn0(xc, X3, xR))) | (xc = X3)) & aReductOfIn0(X3, xa, xR) & aElement0(X3)) | ~ sP15(X2)), inference(nnf_transformation, [], [f73])).
fof(f558, plain, ! [X0, X1] : ~ sP14(X0, X1), inference(resolution, [], [f557, f289])).
fof(f289, plain, ! [X0, X1] : (sP11(sK42(X0, X1)) | ~ sP14(X0, X1)), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ! [X0, X1] : ((sP11(sK42(X0, X1)) & sdtmndtasgtdt0(X0, xR, sK42(X0, X1)) & sP13(sK42(X0, X1), X0) & sdtmndtasgtdt0(X1, xR, sK42(X0, X1)) & sP12(sK42(X0, X1), X1) & aElement0(sK42(X0, X1))) | ~ sP14(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42])], [f141, f142])).
fof(f142, plain, ! [X1, X0] : (? [X2] : (sP11(X2) & sdtmndtasgtdt0(X0, xR, X2) & sP13(X2, X0) & sdtmndtasgtdt0(X1, xR, X2) & sP12(X2, X1) & aElement0(X2)) => (sP11(sK42(X0, X1)) & sdtmndtasgtdt0(X0, xR, sK42(X0, X1)) & sP13(sK42(X0, X1), X0) & sdtmndtasgtdt0(X1, xR, sK42(X0, X1)) & sP12(sK42(X0, X1), X1) & aElement0(sK42(X0, X1)))), introduced(choice_axiom, [])).
fof(f141, plain, ! [X0, X1] : (? [X2] : (sP11(X2) & sdtmndtasgtdt0(X0, xR, X2) & sP13(X2, X0) & sdtmndtasgtdt0(X1, xR, X2) & sP12(X2, X1) & aElement0(X2)) | ~ sP14(X0, X1)), inference(rectify, [], [f140])).
fof(f140, plain, ! [X3, X2] : (? [X4] : (sP11(X4) & sdtmndtasgtdt0(X3, xR, X4) & sP13(X4, X3) & sdtmndtasgtdt0(X2, xR, X4) & sP12(X4, X2) & aElement0(X4)) | ~ sP14(X3, X2)), inference(nnf_transformation, [], [f72])).
fof(f557, plain, ! [X0] : ~ sP11(X0), inference(subsumption_resolution, [], [f556, f298])).
fof(f298, plain, ! [X0] : (aElement0(sK45(X0)) | ~ sP11(X0)), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ! [X0] : ((sdtmndtasgtdt0(xc, xR, sK45(X0)) & sP10(sK45(X0)) & sdtmndtasgtdt0(xb, xR, sK45(X0)) & sP9(sK45(X0)) & aNormalFormOfIn0(sK45(X0), X0, xR) & ! [X2] : ~ aReductOfIn0(X2, sK45(X0), xR) & sdtmndtasgtdt0(X0, xR, sK45(X0)) & sP8(sK45(X0), X0) & aElement0(sK45(X0))) | ~ sP11(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK45])], [f153, f154])).
fof(f154, plain, ! [X0] : (? [X1] : (sdtmndtasgtdt0(xc, xR, X1) & sP10(X1) & sdtmndtasgtdt0(xb, xR, X1) & sP9(X1) & aNormalFormOfIn0(X1, X0, xR) & ! [X2] : ~ aReductOfIn0(X2, X1, xR) & sdtmndtasgtdt0(X0, xR, X1) & sP8(X1, X0) & aElement0(X1)) => (sdtmndtasgtdt0(xc, xR, sK45(X0)) & sP10(sK45(X0)) & sdtmndtasgtdt0(xb, xR, sK45(X0)) & sP9(sK45(X0)) & aNormalFormOfIn0(sK45(X0), X0, xR) & ! [X2] : ~ aReductOfIn0(X2, sK45(X0), xR) & sdtmndtasgtdt0(X0, xR, sK45(X0)) & sP8(sK45(X0), X0) & aElement0(sK45(X0)))), introduced(choice_axiom, [])).
fof(f153, plain, ! [X0] : (? [X1] : (sdtmndtasgtdt0(xc, xR, X1) & sP10(X1) & sdtmndtasgtdt0(xb, xR, X1) & sP9(X1) & aNormalFormOfIn0(X1, X0, xR) & ! [X2] : ~ aReductOfIn0(X2, X1, xR) & sdtmndtasgtdt0(X0, xR, X1) & sP8(X1, X0) & aElement0(X1)) | ~ sP11(X0)), inference(rectify, [], [f152])).
fof(f152, plain, ! [X4] : (? [X5] : (sdtmndtasgtdt0(xc, xR, X5) & sP10(X5) & sdtmndtasgtdt0(xb, xR, X5) & sP9(X5) & aNormalFormOfIn0(X5, X4, xR) & ! [X8] : ~ aReductOfIn0(X8, X5, xR) & sdtmndtasgtdt0(X4, xR, X5) & sP8(X5, X4) & aElement0(X5)) | ~ sP11(X4)), inference(nnf_transformation, [], [f69])).
fof(f556, plain, ! [X0] : (~ aElement0(sK45(X0)) | ~ sP11(X0)), inference(subsumption_resolution, [], [f554, f541])).
fof(f541, plain, ! [X0] : (~ sP18(sK45(X0)) | ~ sP11(X0)), inference(resolution, [], [f304, f264])).
fof(f264, plain, ! [X0] : (~ sdtmndtasgtdt0(xb, xR, X0) | ~ sP18(X0)), inference(cnf_transformation, [], [f127])).
fof(f304, plain, ! [X0] : (sdtmndtasgtdt0(xb, xR, sK45(X0)) | ~ sP11(X0)), inference(cnf_transformation, [], [f155])).
fof(f554, plain, ! [X0] : (sP18(sK45(X0)) | ~ aElement0(sK45(X0)) | ~ sP11(X0)), inference(resolution, [], [f336, f306])).
fof(f306, plain, ! [X0] : (sdtmndtasgtdt0(xc, xR, sK45(X0)) | ~ sP11(X0)), inference(cnf_transformation, [], [f155])).
fof(f336, plain, ! [X0] : (~ sdtmndtasgtdt0(xc, xR, X0) | sP18(X0) | ~ aElement0(X0)), inference(cnf_transformation, [], [f171])).
fof(f371, plain, (sP15(sK38) | ~ spl51_6), inference(avatar_component_clause, [], [f369])).
fof(f369, plain, (spl51_6 <=> sP15(sK38)), introduced(avatar_definition, [new_symbols(naming, [spl51_6])])).
fof(f637, plain, (~ spl51_17 | ~ spl51_22), inference(avatar_contradiction_clause, [], [f636])).
fof(f636, plain, ($false | (~ spl51_17 | ~ spl51_22)), inference(subsumption_resolution, [], [f635, f416])).
fof(f416, plain, sP18(xc), inference(subsumption_resolution, [], [f341, f237])).
fof(f237, plain, aElement0(xc), inference(cnf_transformation, [], [f17])).
fof(f341, plain, (sP18(xc) | ~ aElement0(xc)), inference(equality_resolution, [], [f332])).
fof(f332, plain, ! [X0] : (~ (xc = X0) | sP18(X0) | ~ aElement0(X0)), inference(cnf_transformation, [], [f171])).
fof(f635, plain, (~ sP18(xc) | (~ spl51_17 | ~ spl51_22)), inference(resolution, [], [f569, f424])).
fof(f424, plain, (sdtmndtplgtdt0(xa, xR, xc) | ~ spl51_17), inference(avatar_component_clause, [], [f422])).
fof(f422, plain, (spl51_17 <=> sdtmndtplgtdt0(xa, xR, xc)), introduced(avatar_definition, [new_symbols(naming, [spl51_17])])).
fof(f569, plain, (! [X0] : (~ sdtmndtplgtdt0(xa, xR, X0) | ~ sP18(X0)) | ~ spl51_22), inference(backward_demodulation, [], [f263, f448])).
fof(f448, plain, ((xa = xb) | ~ spl51_22), inference(avatar_component_clause, [], [f446])).
fof(f446, plain, (spl51_22 <=> (xa = xb)), introduced(avatar_definition, [new_symbols(naming, [spl51_22])])).
fof(f263, plain, ! [X0] : (~ sdtmndtplgtdt0(xb, xR, X0) | ~ sP18(X0)), inference(cnf_transformation, [], [f127])).
fof(f621, plain, (~ spl51_16 | ~ spl51_22), inference(avatar_contradiction_clause, [], [f620])).
fof(f620, plain, ($false | (~ spl51_16 | ~ spl51_22)), inference(subsumption_resolution, [], [f605, f577])).
fof(f577, plain, (~ sP18(xa) | ~ spl51_22), inference(backward_demodulation, [], [f340, f448])).
fof(f605, plain, (sP18(xa) | ~ spl51_16), inference(backward_demodulation, [], [f416, f420])).
fof(f465, plain, (spl51_1 | ~ spl51_17 | spl51_5), inference(avatar_split_clause, [], [f321, f365, f422, f347])).
fof(f347, plain, (spl51_1 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl51_1])])).
fof(f365, plain, (spl51_5 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl51_5])])).
fof(f321, plain, (sP16 | ~ sdtmndtplgtdt0(xa, xR, xc) | sP17), inference(cnf_transformation, [], [f171])).
fof(f449, plain, (spl51_22 | spl51_2), inference(avatar_split_clause, [], [f325, f351, f446])).
fof(f325, plain, (sdtmndtplgtdt0(xa, xR, xb) | (xa = xb)), inference(cnf_transformation, [], [f171])).
fof(f425, plain, (spl51_16 | spl51_17), inference(avatar_split_clause, [], [f330, f422, f418])).
fof(f330, plain, (sdtmndtplgtdt0(xa, xR, xc) | (xa = xc)), inference(cnf_transformation, [], [f171])).
fof(f372, plain, (~ spl51_5 | spl51_6), inference(avatar_split_clause, [], [f275, f369, f365])).
fof(f275, plain, (sP15(sK38) | ~ sP16), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ((sP15(sK38) & sdtmndtasgtdt0(sK38, xR, xb) & ((sdtmndtplgtdt0(sK38, xR, xb) & ((sdtmndtplgtdt0(sK39, xR, xb) & aReductOfIn0(sK39, sK38, xR) & aElement0(sK39)) | aReductOfIn0(xb, sK38, xR))) | (xb = sK38)) & aReductOfIn0(sK38, xa, xR) & aElement0(sK38)) | ~ sP16), inference(skolemisation, [status(esa), new_symbols(skolem, [sK38, sK39])], [f131, f133, f132])).
fof(f132, plain, (? [X0] : (sP15(X0) & sdtmndtasgtdt0(X0, xR, xb) & ((sdtmndtplgtdt0(X0, xR, xb) & (? [X1] : (sdtmndtplgtdt0(X1, xR, xb) & aReductOfIn0(X1, X0, xR) & aElement0(X1)) | aReductOfIn0(xb, X0, xR))) | (xb = X0)) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) => (sP15(sK38) & sdtmndtasgtdt0(sK38, xR, xb) & ((sdtmndtplgtdt0(sK38, xR, xb) & (? [X1] : (sdtmndtplgtdt0(X1, xR, xb) & aReductOfIn0(X1, sK38, xR) & aElement0(X1)) | aReductOfIn0(xb, sK38, xR))) | (xb = sK38)) & aReductOfIn0(sK38, xa, xR) & aElement0(sK38))), introduced(choice_axiom, [])).
fof(f133, plain, (? [X1] : (sdtmndtplgtdt0(X1, xR, xb) & aReductOfIn0(X1, sK38, xR) & aElement0(X1)) => (sdtmndtplgtdt0(sK39, xR, xb) & aReductOfIn0(sK39, sK38, xR) & aElement0(sK39))), introduced(choice_axiom, [])).
fof(f131, plain, (? [X0] : (sP15(X0) & sdtmndtasgtdt0(X0, xR, xb) & ((sdtmndtplgtdt0(X0, xR, xb) & (? [X1] : (sdtmndtplgtdt0(X1, xR, xb) & aReductOfIn0(X1, X0, xR) & aElement0(X1)) | aReductOfIn0(xb, X0, xR))) | (xb = X0)) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | ~ sP16), inference(rectify, [], [f130])).
fof(f130, plain, (? [X2] : (sP15(X2) & sdtmndtasgtdt0(X2, xR, xb) & ((sdtmndtplgtdt0(X2, xR, xb) & (? [X13] : (sdtmndtplgtdt0(X13, xR, xb) & aReductOfIn0(X13, X2, xR) & aElement0(X13)) | aReductOfIn0(xb, X2, xR))) | (xb = X2)) & aReductOfIn0(X2, xa, xR) & aElement0(X2)) | ~ sP16), inference(nnf_transformation, [], [f74])).
fof(f354, plain, (~ spl51_1 | ~ spl51_2), inference(avatar_split_clause, [], [f267, f351, f347])).
fof(f267, plain, (~ sdtmndtplgtdt0(xa, xR, xb) | ~ sP17), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ((~ sdtmndtplgtdt0(xa, xR, xb) & ! [X0] : (~ sdtmndtplgtdt0(X0, xR, xb) | ~ aReductOfIn0(X0, xa, xR) | ~ aElement0(X0)) & ~ aReductOfIn0(xb, xa, xR)) | ~ sP17), inference(rectify, [], [f128])).
fof(f128, plain, ((~ sdtmndtplgtdt0(xa, xR, xb) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, xb) | ~ aReductOfIn0(X1, xa, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(xb, xa, xR)) | ~ sP17), inference(nnf_transformation, [], [f75])).