fof(f1649, plain, $false, inference(avatar_sat_refutation, [], [f257, f305, f1053, f1647])).
fof(f1647, plain, (~ spl33_1 | ~ spl33_4), inference(avatar_contradiction_clause, [], [f1646])).
fof(f1646, plain, ($false | (~ spl33_1 | ~ spl33_4)), inference(subsumption_resolution, [], [f1645, f213])).
fof(f213, plain, aElement0(sK29), inference(cnf_transformation, [], [f121])).
fof(f121, plain, (~ isConfluent0(xR) & (! [X3] : ((~ sdtmndtasgtdt0(sK30, xR, X3) & ~ sdtmndtplgtdt0(sK30, xR, X3) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X3) | ~ aReductOfIn0(X4, sK30, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X3, sK30, xR) & ~ (sK30 = X3)) | sP8(X3, sK29) | ~ aElement0(X3)) & sdtmndtasgtdt0(sK28, xR, sK30) & ((sdtmndtplgtdt0(sK28, xR, sK30) & ((sdtmndtplgtdt0(sK31, xR, sK30) & aReductOfIn0(sK31, sK28, xR) & aElement0(sK31)) | aReductOfIn0(sK30, sK28, xR))) | (sK28 = sK30)) & sdtmndtasgtdt0(sK28, xR, sK29) & ((sdtmndtplgtdt0(sK28, xR, sK29) & ((sdtmndtplgtdt0(sK32, xR, sK29) & aReductOfIn0(sK32, sK28, xR) & aElement0(sK32)) | aReductOfIn0(sK29, sK28, xR))) | (sK28 = sK29)) & aElement0(sK30) & aElement0(sK29) & aElement0(sK28))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK28, sK29, sK30, sK31, sK32])], [f117, f120, f119, f118])).
fof(f118, plain, (? [X0, X1, X2] : (! [X3] : ((~ sdtmndtasgtdt0(X2, xR, X3) & ~ sdtmndtplgtdt0(X2, xR, X3) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X3) | ~ aReductOfIn0(X4, X2, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X3, X2, xR) & ~ (X2 = X3)) | sP8(X3, X1) | ~ aElement0(X3)) & sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X5] : (sdtmndtplgtdt0(X5, xR, X2) & aReductOfIn0(X5, X0, xR) & aElement0(X5)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X0, xR, X1) & ((sdtmndtplgtdt0(X0, xR, X1) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X1) & aReductOfIn0(X6, X0, xR) & aElement0(X6)) | aReductOfIn0(X1, X0, xR))) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0)) => (! [X3] : ((~ sdtmndtasgtdt0(sK30, xR, X3) & ~ sdtmndtplgtdt0(sK30, xR, X3) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X3) | ~ aReductOfIn0(X4, sK30, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X3, sK30, xR) & ~ (sK30 = X3)) | sP8(X3, sK29) | ~ aElement0(X3)) & sdtmndtasgtdt0(sK28, xR, sK30) & ((sdtmndtplgtdt0(sK28, xR, sK30) & (? [X5] : (sdtmndtplgtdt0(X5, xR, sK30) & aReductOfIn0(X5, sK28, xR) & aElement0(X5)) | aReductOfIn0(sK30, sK28, xR))) | (sK28 = sK30)) & sdtmndtasgtdt0(sK28, xR, sK29) & ((sdtmndtplgtdt0(sK28, xR, sK29) & (? [X6] : (sdtmndtplgtdt0(X6, xR, sK29) & aReductOfIn0(X6, sK28, xR) & aElement0(X6)) | aReductOfIn0(sK29, sK28, xR))) | (sK28 = sK29)) & aElement0(sK30) & aElement0(sK29) & aElement0(sK28))), introduced(choice_axiom, [])).
fof(f119, plain, (? [X5] : (sdtmndtplgtdt0(X5, xR, sK30) & aReductOfIn0(X5, sK28, xR) & aElement0(X5)) => (sdtmndtplgtdt0(sK31, xR, sK30) & aReductOfIn0(sK31, sK28, xR) & aElement0(sK31))), introduced(choice_axiom, [])).
fof(f120, plain, (? [X6] : (sdtmndtplgtdt0(X6, xR, sK29) & aReductOfIn0(X6, sK28, xR) & aElement0(X6)) => (sdtmndtplgtdt0(sK32, xR, sK29) & aReductOfIn0(sK32, sK28, xR) & aElement0(sK32))), introduced(choice_axiom, [])).
fof(f117, plain, (~ isConfluent0(xR) & ? [X0, X1, X2] : (! [X3] : ((~ sdtmndtasgtdt0(X2, xR, X3) & ~ sdtmndtplgtdt0(X2, xR, X3) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X3) | ~ aReductOfIn0(X4, X2, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X3, X2, xR) & ~ (X2 = X3)) | sP8(X3, X1) | ~ aElement0(X3)) & sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X5] : (sdtmndtplgtdt0(X5, xR, X2) & aReductOfIn0(X5, X0, xR) & aElement0(X5)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X0, xR, X1) & ((sdtmndtplgtdt0(X0, xR, X1) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X1) & aReductOfIn0(X6, X0, xR) & aElement0(X6)) | aReductOfIn0(X1, X0, xR))) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0))), inference(rectify, [], [f66])).
fof(f66, plain, (~ isConfluent0(xR) & ? [X0, X1, X2] : (! [X5] : ((~ sdtmndtasgtdt0(X2, xR, X5) & ~ sdtmndtplgtdt0(X2, xR, X5) & ! [X6] : (~ sdtmndtplgtdt0(X6, xR, X5) | ~ aReductOfIn0(X6, X2, xR) | ~ aElement0(X6)) & ~ aReductOfIn0(X5, X2, xR) & ~ (X2 = X5)) | sP8(X5, X1) | ~ aElement0(X5)) & sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X0, xR, X1) & ((sdtmndtplgtdt0(X0, xR, X1) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X1) & aReductOfIn0(X4, X0, xR) & aElement0(X4)) | aReductOfIn0(X1, X0, xR))) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0))), inference(definition_folding, [], [f52, e65])).
fof(f65, plain, ! [X5, X1] : ((~ sdtmndtasgtdt0(X1, xR, X5) & ~ sdtmndtplgtdt0(X1, xR, X5) & ! [X7] : (~ sdtmndtplgtdt0(X7, xR, X5) | ~ aReductOfIn0(X7, X1, xR) | ~ aElement0(X7)) & ~ aReductOfIn0(X5, X1, xR) & ~ (X1 = X5)) | ~ sP8(X5, X1)), inference(usedef, [], [e65])).
fof(e65, plain, ! [X5, X1] : (sP8(X5, X1) <=> (~ sdtmndtasgtdt0(X1, xR, X5) & ~ sdtmndtplgtdt0(X1, xR, X5) & ! [X7] : (~ sdtmndtplgtdt0(X7, xR, X5) | ~ aReductOfIn0(X7, X1, xR) | ~ aElement0(X7)) & ~ aReductOfIn0(X5, X1, xR) & ~ (X1 = X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f52, plain, (~ isConfluent0(xR) & ? [X0, X1, X2] : (! [X5] : ((~ sdtmndtasgtdt0(X2, xR, X5) & ~ sdtmndtplgtdt0(X2, xR, X5) & ! [X6] : (~ sdtmndtplgtdt0(X6, xR, X5) | ~ aReductOfIn0(X6, X2, xR) | ~ aElement0(X6)) & ~ aReductOfIn0(X5, X2, xR) & ~ (X2 = X5)) | (~ sdtmndtasgtdt0(X1, xR, X5) & ~ sdtmndtplgtdt0(X1, xR, X5) & ! [X7] : (~ sdtmndtplgtdt0(X7, xR, X5) | ~ aReductOfIn0(X7, X1, xR) | ~ aElement0(X7)) & ~ aReductOfIn0(X5, X1, xR) & ~ (X1 = X5)) | ~ aElement0(X5)) & sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X0, xR, X1) & ((sdtmndtplgtdt0(X0, xR, X1) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X1) & aReductOfIn0(X4, X0, xR) & aElement0(X4)) | aReductOfIn0(X1, X0, xR))) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0))), inference(flattening, [], [f51])).
fof(f51, plain, (~ isConfluent0(xR) & ? [X0, X1, X2] : (! [X5] : ((~ sdtmndtasgtdt0(X2, xR, X5) & ~ sdtmndtplgtdt0(X2, xR, X5) & ! [X6] : (~ sdtmndtplgtdt0(X6, xR, X5) | ~ aReductOfIn0(X6, X2, xR) | ~ aElement0(X6)) & ~ aReductOfIn0(X5, X2, xR) & ~ (X2 = X5)) | (~ sdtmndtasgtdt0(X1, xR, X5) & ~ sdtmndtplgtdt0(X1, xR, X5) & ! [X7] : (~ sdtmndtplgtdt0(X7, xR, X5) | ~ aReductOfIn0(X7, X1, xR) | ~ aElement0(X7)) & ~ aReductOfIn0(X5, X1, xR) & ~ (X1 = X5)) | ~ aElement0(X5)) & (sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X0, xR, X1) & ((sdtmndtplgtdt0(X0, xR, X1) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X1) & aReductOfIn0(X4, X0, xR) & aElement0(X4)) | aReductOfIn0(X1, X0, xR))) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0)))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ~ (isConfluent0(xR) | ! [X0, X1, X2] : ((sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X0, xR, X1) & ((sdtmndtplgtdt0(X0, xR, X1) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X1) & aReductOfIn0(X4, X0, xR) & aElement0(X4)) | aReductOfIn0(X1, X0, xR))) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0)) => ? [X5] : ((sdtmndtasgtdt0(X2, xR, X5) | sdtmndtplgtdt0(X2, xR, X5) | ? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR) | (X2 = X5)) & (sdtmndtasgtdt0(X1, xR, X5) | sdtmndtplgtdt0(X1, xR, X5) | ? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR) | (X1 = X5)) & aElement0(X5)))), inference(rectify, [], [f19])).
fof(f19, plain, ~ (isConfluent0(xR) | ! [X0, X1, X2] : ((sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X0, xR, X1) & ((sdtmndtplgtdt0(X0, xR, X1) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X1) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X1, X0, xR))) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0)) => ? [X3] : ((sdtmndtasgtdt0(X2, xR, X3) | sdtmndtplgtdt0(X2, xR, X3) | ? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X2, xR) & aElement0(X4)) | aReductOfIn0(X3, X2, xR) | (X2 = X3)) & (sdtmndtasgtdt0(X1, xR, X3) | sdtmndtplgtdt0(X1, xR, X3) | ? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X1, xR) & aElement0(X4)) | aReductOfIn0(X3, X1, xR) | (X1 = X3)) & aElement0(X3)))), inference(negated_conjecture, [], [f18])).
fof(f18, plain, ~ (isConfluent0(xR) | ! [X0, X1, X2] : ((sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X0, xR, X1) & ((sdtmndtplgtdt0(X0, xR, X1) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X1) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X1, X0, xR))) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0)) => ? [X3] : ((sdtmndtasgtdt0(X2, xR, X3) | sdtmndtplgtdt0(X2, xR, X3) | ? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X2, xR) & aElement0(X4)) | aReductOfIn0(X3, X2, xR) | (X2 = X3)) & (sdtmndtasgtdt0(X1, xR, X3) | sdtmndtplgtdt0(X1, xR, X3) | ? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X1, xR) & aElement0(X4)) | aReductOfIn0(X3, X1, xR) | (X1 = X3)) & aElement0(X3)))), file('/home/ubuntu/library/tptp/Problems/COM/COM023+4.p', m__)).
fof(f1645, plain, (~ aElement0(sK29) | (~ spl33_1 | ~ spl33_4)), inference(subsumption_resolution, [], [f1640, f488])).
fof(f488, plain, ~ sP6(sK30, sK29), inference(duplicate_literal_removal, [], [f487])).
fof(f487, plain, (~ sP6(sK30, sK29) | ~ sP6(sK30, sK29)), inference(resolution, [], [f352, f359])).
fof(f359, plain, ! [X4] : (sP8(sK25(sK30, X4), sK29) | ~ sP6(sK30, X4)), inference(subsumption_resolution, [], [f358, f190])).
fof(f190, plain, ! [X0, X1] : (aElement0(sK25(X0, X1)) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ! [X0, X1] : ((sdtmndtasgtdt0(X0, xR, sK25(X0, X1)) & ((sdtmndtplgtdt0(X0, xR, sK25(X0, X1)) & ((sdtmndtplgtdt0(sK26(X0, X1), xR, sK25(X0, X1)) & aReductOfIn0(sK26(X0, X1), X0, xR) & aElement0(sK26(X0, X1))) | aReductOfIn0(sK25(X0, X1), X0, xR))) | (sK25(X0, X1) = X0)) & sdtmndtasgtdt0(X1, xR, sK25(X0, X1)) & sP5(sK25(X0, X1), X1) & aElement0(sK25(X0, X1))) | ~ sP6(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK25, sK26])], [f107, f109, f108])).
fof(f108, plain, ! [X1, X0] : (? [X2] : (sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X1, xR, X2) & sP5(X2, X1) & aElement0(X2)) => (sdtmndtasgtdt0(X0, xR, sK25(X0, X1)) & ((sdtmndtplgtdt0(X0, xR, sK25(X0, X1)) & (? [X3] : (sdtmndtplgtdt0(X3, xR, sK25(X0, X1)) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(sK25(X0, X1), X0, xR))) | (sK25(X0, X1) = X0)) & sdtmndtasgtdt0(X1, xR, sK25(X0, X1)) & sP5(sK25(X0, X1), X1) & aElement0(sK25(X0, X1)))), introduced(choice_axiom, [])).
fof(f109, plain, ! [X1, X0] : (? [X3] : (sdtmndtplgtdt0(X3, xR, sK25(X0, X1)) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) => (sdtmndtplgtdt0(sK26(X0, X1), xR, sK25(X0, X1)) & aReductOfIn0(sK26(X0, X1), X0, xR) & aElement0(sK26(X0, X1)))), introduced(choice_axiom, [])).
fof(f107, plain, ! [X0, X1] : (? [X2] : (sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X1, xR, X2) & sP5(X2, X1) & aElement0(X2)) | ~ sP6(X0, X1)), inference(rectify, [], [f106])).
fof(f106, plain, ! [X2, X1] : (? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & sP5(X5, X1) & aElement0(X5)) | ~ sP6(X2, X1)), inference(nnf_transformation, [], [f62])).
fof(f62, plain, ! [X2, X1] : (? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & sP5(X5, X1) & aElement0(X5)) | ~ sP6(X2, X1)), inference(usedef, [], [e62])).
fof(e62, plain, ! [X2, X1] : (sP6(X2, X1) <=> ? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & sP5(X5, X1) & aElement0(X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f358, plain, ! [X4] : (~ sP6(sK30, X4) | sP8(sK25(sK30, X4), sK29) | ~ aElement0(sK25(sK30, X4))), inference(resolution, [], [f197, f229])).
fof(f229, plain, ! [X3] : (~ sdtmndtasgtdt0(sK30, xR, X3) | sP8(X3, sK29) | ~ aElement0(X3)), inference(cnf_transformation, [], [f121])).
fof(f197, plain, ! [X0, X1] : (sdtmndtasgtdt0(X0, xR, sK25(X0, X1)) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f110])).
fof(f352, plain, ! [X0, X1] : (~ sP8(sK25(X0, X1), X1) | ~ sP6(X0, X1)), inference(resolution, [], [f192, f211])).
fof(f211, plain, ! [X0, X1] : (~ sdtmndtasgtdt0(X1, xR, X0) | ~ sP8(X0, X1)), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ! [X0, X1] : ((~ sdtmndtasgtdt0(X1, xR, X0) & ~ sdtmndtplgtdt0(X1, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, X1, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, X1, xR) & ~ (X0 = X1)) | ~ sP8(X0, X1)), inference(rectify, [], [f115])).
fof(f115, plain, ! [X5, X1] : ((~ sdtmndtasgtdt0(X1, xR, X5) & ~ sdtmndtplgtdt0(X1, xR, X5) & ! [X7] : (~ sdtmndtplgtdt0(X7, xR, X5) | ~ aReductOfIn0(X7, X1, xR) | ~ aElement0(X7)) & ~ aReductOfIn0(X5, X1, xR) & ~ (X1 = X5)) | ~ sP8(X5, X1)), inference(nnf_transformation, [], [f65])).
fof(f192, plain, ! [X0, X1] : (sdtmndtasgtdt0(X1, xR, sK25(X0, X1)) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f110])).
fof(f1640, plain, (sP6(sK30, sK29) | ~ aElement0(sK29) | (~ spl33_1 | ~ spl33_4)), inference(resolution, [], [f576, f328])).
fof(f328, plain, ~ sP7(sK29, sK28), inference(resolution, [], [f219, f189])).
fof(f189, plain, ! [X0, X1] : (~ sdtmndtasgtdt0(X1, xR, X0) | ~ sP7(X0, X1)), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ! [X0, X1] : ((~ sdtmndtasgtdt0(X1, xR, X0) & ~ sdtmndtplgtdt0(X1, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, X1, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, X1, xR) & ~ (X0 = X1)) | ~ sP7(X0, X1)), inference(rectify, [], [f104])).
fof(f104, plain, ! [X1, X0] : ((~ sdtmndtasgtdt0(X0, xR, X1) & ~ sdtmndtplgtdt0(X0, xR, X1) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X1) | ~ aReductOfIn0(X4, X0, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X1, X0, xR) & ~ (X0 = X1)) | ~ sP7(X1, X0)), inference(nnf_transformation, [], [f63])).
fof(f63, plain, ! [X1, X0] : ((~ sdtmndtasgtdt0(X0, xR, X1) & ~ sdtmndtplgtdt0(X0, xR, X1) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X1) | ~ aReductOfIn0(X4, X0, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X1, X0, xR) & ~ (X0 = X1)) | ~ sP7(X1, X0)), inference(usedef, [], [e63])).
fof(e63, plain, ! [X1, X0] : (sP7(X1, X0) <=> (~ sdtmndtasgtdt0(X0, xR, X1) & ~ sdtmndtplgtdt0(X0, xR, X1) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X1) | ~ aReductOfIn0(X4, X0, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X1, X0, xR) & ~ (X0 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f219, plain, sdtmndtasgtdt0(sK28, xR, sK29), inference(cnf_transformation, [], [f121])).
fof(f576, plain, (! [X0] : (sP7(X0, sK28) | sP6(sK30, X0) | ~ aElement0(X0)) | (~ spl33_1 | ~ spl33_4)), inference(subsumption_resolution, [], [f575, f212])).
fof(f212, plain, aElement0(sK28), inference(cnf_transformation, [], [f121])).
fof(f575, plain, (! [X0] : (sP6(sK30, X0) | sP7(X0, sK28) | ~ aElement0(X0) | ~ aElement0(sK28)) | (~ spl33_1 | ~ spl33_4)), inference(subsumption_resolution, [], [f559, f242])).
fof(f242, plain, (aElement0(sK30) | ~ spl33_1), inference(avatar_component_clause, [], [f241])).
fof(f241, plain, (spl33_1 <=> aElement0(sK30)), introduced(avatar_definition, [new_symbols(naming, [spl33_1])])).
fof(f559, plain, (! [X0] : (sP6(sK30, X0) | sP7(X0, sK28) | ~ aElement0(sK30) | ~ aElement0(X0) | ~ aElement0(sK28)) | ~ spl33_4), inference(resolution, [], [f205, f256])).
fof(f256, plain, (sdtmndtplgtdt0(sK28, xR, sK30) | ~ spl33_4), inference(avatar_component_clause, [], [f254])).
fof(f254, plain, (spl33_4 <=> sdtmndtplgtdt0(sK28, xR, sK30)), introduced(avatar_definition, [new_symbols(naming, [spl33_4])])).
fof(f205, plain, ! [X2, X0, X1] : (~ sdtmndtplgtdt0(X0, xR, X2) | sP6(X2, X1) | sP7(X1, X0) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f64])).
fof(f64, plain, ! [X0, X1, X2] : (sP6(X2, X1) | (~ sdtmndtasgtdt0(X0, xR, X2) & ~ sdtmndtplgtdt0(X0, xR, X2) & ! [X3] : (~ sdtmndtplgtdt0(X3, xR, X2) | ~ aReductOfIn0(X3, X0, xR) | ~ aElement0(X3)) & ~ aReductOfIn0(X2, X0, xR) & ~ (X0 = X2)) | sP7(X1, X0) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0)), inference(definition_folding, [], [f50, e63, e62, e61])).
fof(f61, plain, ! [X5, X1] : ((sdtmndtplgtdt0(X1, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR))) | (X1 = X5) | ~ sP5(X5, X1)), inference(usedef, [], [e61])).
fof(e61, plain, ! [X5, X1] : (sP5(X5, X1) <=> ((sdtmndtplgtdt0(X1, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR))) | (X1 = X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f50, plain, ! [X0, X1, X2] : (? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & ((sdtmndtplgtdt0(X1, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR))) | (X1 = X5)) & aElement0(X5)) | (~ sdtmndtasgtdt0(X0, xR, X2) & ~ sdtmndtplgtdt0(X0, xR, X2) & ! [X3] : (~ sdtmndtplgtdt0(X3, xR, X2) | ~ aReductOfIn0(X3, X0, xR) | ~ aElement0(X3)) & ~ aReductOfIn0(X2, X0, xR) & ~ (X0 = X2)) | (~ sdtmndtasgtdt0(X0, xR, X1) & ~ sdtmndtplgtdt0(X0, xR, X1) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X1) | ~ aReductOfIn0(X4, X0, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X1, X0, xR) & ~ (X0 = X1)) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0)), inference(flattening, [], [f49])).
fof(f49, plain, ! [X0, X1, X2] : (? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & ((sdtmndtplgtdt0(X1, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR))) | (X1 = X5)) & aElement0(X5)) | ((~ sdtmndtasgtdt0(X0, xR, X2) & ~ sdtmndtplgtdt0(X0, xR, X2) & ! [X3] : (~ sdtmndtplgtdt0(X3, xR, X2) | ~ aReductOfIn0(X3, X0, xR) | ~ aElement0(X3)) & ~ aReductOfIn0(X2, X0, xR) & ~ (X0 = X2)) | (~ sdtmndtasgtdt0(X0, xR, X1) & ~ sdtmndtplgtdt0(X0, xR, X1) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X1) | ~ aReductOfIn0(X4, X0, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X1, X0, xR) & ~ (X0 = X1)) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ! [X0, X1, X2] : (((sdtmndtasgtdt0(X0, xR, X2) | sdtmndtplgtdt0(X0, xR, X2) | ? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR) | (X0 = X2)) & (sdtmndtasgtdt0(X0, xR, X1) | sdtmndtplgtdt0(X0, xR, X1) | ? [X4] : (sdtmndtplgtdt0(X4, xR, X1) & aReductOfIn0(X4, X0, xR) & aElement0(X4)) | aReductOfIn0(X1, X0, xR) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0)) => ? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & ((sdtmndtplgtdt0(X1, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR))) | (X1 = X5)) & aElement0(X5))), inference(rectify, [], [f17])).
fof(f17, plain, ! [X0, X1, X2] : (((sdtmndtasgtdt0(X0, xR, X2) | sdtmndtplgtdt0(X0, xR, X2) | ? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR) | (X0 = X2)) & (sdtmndtasgtdt0(X0, xR, X1) | sdtmndtplgtdt0(X0, xR, X1) | ? [X3] : (sdtmndtplgtdt0(X3, xR, X1) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X1, X0, xR) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0)) => ? [X3] : (sdtmndtasgtdt0(X2, xR, X3) & ((sdtmndtplgtdt0(X2, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X2, xR) & aElement0(X4)) | aReductOfIn0(X3, X2, xR))) | (X2 = X3)) & sdtmndtasgtdt0(X1, xR, X3) & ((sdtmndtplgtdt0(X1, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X1, xR) & aElement0(X4)) | aReductOfIn0(X3, X1, xR))) | (X1 = X3)) & aElement0(X3))), file('/home/ubuntu/library/tptp/Problems/COM/COM023+4.p', m__715)).
fof(f1053, plain, ~ spl33_3, inference(avatar_contradiction_clause, [], [f1052])).
fof(f1052, plain, ($false | ~ spl33_3), inference(subsumption_resolution, [], [f1025, f389])).
fof(f389, plain, sP6(sK28, sK29), inference(subsumption_resolution, [], [f388, f213])).
fof(f388, plain, (sP6(sK28, sK29) | ~ aElement0(sK29)), inference(subsumption_resolution, [], [f363, f212])).
fof(f363, plain, (sP6(sK28, sK29) | ~ aElement0(sK28) | ~ aElement0(sK29)), inference(resolution, [], [f236, f328])).
fof(f236, plain, ! [X2, X1] : (sP7(X1, X2) | sP6(X2, X1) | ~ aElement0(X2) | ~ aElement0(X1)), inference(duplicate_literal_removal, [], [f233])).
fof(f233, plain, ! [X2, X1] : (sP6(X2, X1) | sP7(X1, X2) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X2)), inference(equality_resolution, [], [f202])).
fof(f202, plain, ! [X2, X0, X1] : (sP6(X2, X1) | ~ (X0 = X2) | sP7(X1, X0) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f64])).
fof(f1025, plain, (~ sP6(sK28, sK29) | ~ spl33_3), inference(backward_demodulation, [], [f488, f252])).
fof(f252, plain, ((sK28 = sK30) | ~ spl33_3), inference(avatar_component_clause, [], [f250])).
fof(f250, plain, (spl33_3 <=> (sK28 = sK30)), introduced(avatar_definition, [new_symbols(naming, [spl33_3])])).
fof(f305, plain, spl33_1, inference(avatar_split_clause, [], [f214, f241])).
fof(f214, plain, aElement0(sK30), inference(cnf_transformation, [], [f121])).
fof(f257, plain, (spl33_3 | spl33_4), inference(avatar_split_clause, [], [f223, f254, f250])).
fof(f223, plain, (sdtmndtplgtdt0(sK28, xR, sK30) | (sK28 = sK30)), inference(cnf_transformation, [], [f121])).