fof(f370, plain, $false, inference(avatar_sat_refutation, [], [f185, f324, f332, f367, f369])).
fof(f369, plain, ~ spl23_3, inference(avatar_contradiction_clause, [], [f368])).
fof(f368, plain, ($false | ~ spl23_3), inference(subsumption_resolution, [], [f184, f251])).
fof(f251, plain, ! [X0, X1] : ~ sP4(X0, X1), inference(subsumption_resolution, [], [f250, f153])).
fof(f153, plain, ! [X0, X1] : (aElement0(sK19(X0, X1)) | ~ sP4(X0, X1)), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ! [X0, X1] : (((sdtmndtasgtdt0(xc, xR, sK20(X0, X1)) & sdtmndtasgtdt0(xb, xR, sK20(X0, X1)) & aNormalFormOfIn0(sK20(X0, X1), sK19(X0, X1), xR)) & sdtmndtasgtdt0(X0, xR, sK19(X0, X1)) & sdtmndtasgtdt0(X1, xR, sK19(X0, X1)) & aElement0(sK19(X0, X1))) | ~ sP4(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK19, sK20])], [f91, f93, f92])).
fof(f92, plain, ! [X1, X0] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, X2, xR)) & sdtmndtasgtdt0(X0, xR, X2) & sdtmndtasgtdt0(X1, xR, X2) & aElement0(X2)) => (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, sK19(X0, X1), xR)) & sdtmndtasgtdt0(X0, xR, sK19(X0, X1)) & sdtmndtasgtdt0(X1, xR, sK19(X0, X1)) & aElement0(sK19(X0, X1)))), introduced(choice_axiom, [])).
fof(f93, plain, ! [X1, X0] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, sK19(X0, X1), xR)) => (sdtmndtasgtdt0(xc, xR, sK20(X0, X1)) & sdtmndtasgtdt0(xb, xR, sK20(X0, X1)) & aNormalFormOfIn0(sK20(X0, X1), sK19(X0, X1), xR))), introduced(choice_axiom, [])).
fof(f91, plain, ! [X0, X1] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, X2, xR)) & sdtmndtasgtdt0(X0, xR, X2) & sdtmndtasgtdt0(X1, xR, X2) & aElement0(X2)) | ~ sP4(X0, X1)), inference(rectify, [], [f90])).
fof(f90, plain, ! [X1, X0] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, X2, xR)) & sdtmndtasgtdt0(X1, xR, X2) & sdtmndtasgtdt0(X0, xR, X2) & aElement0(X2)) | ~ sP4(X1, X0)), inference(nnf_transformation, [], [f56])).
fof(f56, plain, ! [X1, X0] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, X2, xR)) & sdtmndtasgtdt0(X1, xR, X2) & sdtmndtasgtdt0(X0, xR, X2) & aElement0(X2)) | ~ sP4(X1, X0)), inference(usedef, [], [e56])).
fof(e56, plain, ! [X1, X0] : (sP4(X1, X0) <=> ? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, X2, xR)) & sdtmndtasgtdt0(X1, xR, X2) & sdtmndtasgtdt0(X0, xR, X2) & aElement0(X2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f250, plain, ! [X0, X1] : (~ sP4(X0, X1) | ~ aElement0(sK19(X0, X1))), inference(subsumption_resolution, [], [f249, f239])).
fof(f239, plain, ! [X0, X1] : (~ aElement0(sK20(X0, X1)) | ~ sP4(X0, X1)), inference(subsumption_resolution, [], [f234, f157])).
fof(f157, plain, ! [X0, X1] : (sdtmndtasgtdt0(xb, xR, sK20(X0, X1)) | ~ sP4(X0, X1)), inference(cnf_transformation, [], [f94])).
fof(f234, plain, ! [X0, X1] : (~ sdtmndtasgtdt0(xb, xR, sK20(X0, X1)) | ~ aElement0(sK20(X0, X1)) | ~ sP4(X0, X1)), inference(resolution, [], [f168, f158])).
fof(f158, plain, ! [X0, X1] : (sdtmndtasgtdt0(xc, xR, sK20(X0, X1)) | ~ sP4(X0, X1)), inference(cnf_transformation, [], [f94])).
fof(f168, plain, ! [X0] : (~ sdtmndtasgtdt0(xc, xR, X0) | ~ sdtmndtasgtdt0(xb, xR, X0) | ~ aElement0(X0)), inference(cnf_transformation, [], [f98])).
fof(f98, plain, (! [X0] : (~ sdtmndtasgtdt0(xc, xR, X0) | ~ sdtmndtasgtdt0(xb, xR, X0) | ~ aElement0(X0)) & sdtmndtasgtdt0(xa, xR, xc) & sdtmndtasgtdt0(xa, xR, xb) & (((sP4(sK22, sK21) & sdtmndtasgtdt0(sK22, xR, xc) & aReductOfIn0(sK22, xa, xR) & aElement0(sK22)) & sdtmndtasgtdt0(sK21, xR, xb) & aReductOfIn0(sK21, xa, xR) & aElement0(sK21)) | ~ sdtmndtplgtdt0(xa, xR, xc) | ~ sdtmndtplgtdt0(xa, xR, xb))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21, sK22])], [f95, f97, f96])).
fof(f96, plain, (? [X1] : (? [X2] : (sP4(X2, X1) & sdtmndtasgtdt0(X2, xR, xc) & aReductOfIn0(X2, xa, xR) & aElement0(X2)) & sdtmndtasgtdt0(X1, xR, xb) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) => (? [X2] : (sP4(X2, sK21) & sdtmndtasgtdt0(X2, xR, xc) & aReductOfIn0(X2, xa, xR) & aElement0(X2)) & sdtmndtasgtdt0(sK21, xR, xb) & aReductOfIn0(sK21, xa, xR) & aElement0(sK21))), introduced(choice_axiom, [])).
fof(f97, plain, (? [X2] : (sP4(X2, sK21) & sdtmndtasgtdt0(X2, xR, xc) & aReductOfIn0(X2, xa, xR) & aElement0(X2)) => (sP4(sK22, sK21) & sdtmndtasgtdt0(sK22, xR, xc) & aReductOfIn0(sK22, xa, xR) & aElement0(sK22))), introduced(choice_axiom, [])).
fof(f95, plain, (! [X0] : (~ sdtmndtasgtdt0(xc, xR, X0) | ~ sdtmndtasgtdt0(xb, xR, X0) | ~ aElement0(X0)) & sdtmndtasgtdt0(xa, xR, xc) & sdtmndtasgtdt0(xa, xR, xb) & (? [X1] : (? [X2] : (sP4(X2, X1) & sdtmndtasgtdt0(X2, xR, xc) & aReductOfIn0(X2, xa, xR) & aElement0(X2)) & sdtmndtasgtdt0(X1, xR, xb) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) | ~ sdtmndtplgtdt0(xa, xR, xc) | ~ sdtmndtplgtdt0(xa, xR, xb))), inference(rectify, [], [f57])).
fof(f57, plain, (! [X4] : (~ sdtmndtasgtdt0(xc, xR, X4) | ~ sdtmndtasgtdt0(xb, xR, X4) | ~ aElement0(X4)) & sdtmndtasgtdt0(xa, xR, xc) & sdtmndtasgtdt0(xa, xR, xb) & (? [X0] : (? [X1] : (sP4(X1, X0) & sdtmndtasgtdt0(X1, xR, xc) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) & sdtmndtasgtdt0(X0, xR, xb) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | ~ sdtmndtplgtdt0(xa, xR, xc) | ~ sdtmndtplgtdt0(xa, xR, xb))), inference(definition_folding, [], [f49, e56])).
fof(f49, plain, (! [X4] : (~ sdtmndtasgtdt0(xc, xR, X4) | ~ sdtmndtasgtdt0(xb, xR, X4) | ~ aElement0(X4)) & sdtmndtasgtdt0(xa, xR, xc) & sdtmndtasgtdt0(xa, xR, xb) & (? [X0] : (? [X1] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, X2, xR)) & sdtmndtasgtdt0(X1, xR, X2) & sdtmndtasgtdt0(X0, xR, X2) & aElement0(X2)) & sdtmndtasgtdt0(X1, xR, xc) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) & sdtmndtasgtdt0(X0, xR, xb) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | ~ sdtmndtplgtdt0(xa, xR, xc) | ~ sdtmndtplgtdt0(xa, xR, xb))), inference(flattening, [], [f48])).
fof(f48, plain, ((! [X4] : (~ sdtmndtasgtdt0(xc, xR, X4) | ~ sdtmndtasgtdt0(xb, xR, X4) | ~ aElement0(X4)) & (sdtmndtasgtdt0(xa, xR, xc) & sdtmndtasgtdt0(xa, xR, xb))) & (? [X0] : (? [X1] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, X2, xR)) & sdtmndtasgtdt0(X1, xR, X2) & sdtmndtasgtdt0(X0, xR, X2) & aElement0(X2)) & sdtmndtasgtdt0(X1, xR, xc) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) & sdtmndtasgtdt0(X0, xR, xb) & aReductOfIn0(X0, xa, xR) & aElement0(X0)) | (~ sdtmndtplgtdt0(xa, xR, xc) | ~ sdtmndtplgtdt0(xa, xR, xb)))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ~ (((sdtmndtplgtdt0(xa, xR, xc) & sdtmndtplgtdt0(xa, xR, xb)) => ? [X0] : (? [X1] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, X2, xR)) & sdtmndtasgtdt0(X1, xR, X2) & sdtmndtasgtdt0(X0, xR, X2) & aElement0(X2)) & sdtmndtasgtdt0(X1, xR, xc) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) & sdtmndtasgtdt0(X0, xR, xb) & aReductOfIn0(X0, xa, xR) & aElement0(X0))) => ((sdtmndtasgtdt0(xa, xR, xc) & sdtmndtasgtdt0(xa, xR, xb)) => ? [X4] : (sdtmndtasgtdt0(xc, xR, X4) & sdtmndtasgtdt0(xb, xR, X4) & aElement0(X4)))), inference(rectify, [], [f20])).
fof(f20, plain, ~ (((sdtmndtplgtdt0(xa, xR, xc) & sdtmndtplgtdt0(xa, xR, xb)) => ? [X0] : (? [X1] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, X2, xR)) & sdtmndtasgtdt0(X1, xR, X2) & sdtmndtasgtdt0(X0, xR, X2) & aElement0(X2)) & sdtmndtasgtdt0(X1, xR, xc) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) & sdtmndtasgtdt0(X0, xR, xb) & aReductOfIn0(X0, xa, xR) & aElement0(X0))) => ((sdtmndtasgtdt0(xa, xR, xc) & sdtmndtasgtdt0(xa, xR, xb)) => ? [X0] : (sdtmndtasgtdt0(xc, xR, X0) & sdtmndtasgtdt0(xb, xR, X0) & aElement0(X0)))), inference(negated_conjecture, [], [f19])).
fof(f19, plain, ~ (((sdtmndtplgtdt0(xa, xR, xc) & sdtmndtplgtdt0(xa, xR, xb)) => ? [X0] : (? [X1] : (? [X2] : (? [X3] : (sdtmndtasgtdt0(xc, xR, X3) & sdtmndtasgtdt0(xb, xR, X3) & aNormalFormOfIn0(X3, X2, xR)) & sdtmndtasgtdt0(X1, xR, X2) & sdtmndtasgtdt0(X0, xR, X2) & aElement0(X2)) & sdtmndtasgtdt0(X1, xR, xc) & aReductOfIn0(X1, xa, xR) & aElement0(X1)) & sdtmndtasgtdt0(X0, xR, xb) & aReductOfIn0(X0, xa, xR) & aElement0(X0))) => ((sdtmndtasgtdt0(xa, xR, xc) & sdtmndtasgtdt0(xa, xR, xb)) => ? [X0] : (sdtmndtasgtdt0(xc, xR, X0) & sdtmndtasgtdt0(xb, xR, X0) & aElement0(X0)))), file('/home/ubuntu/library/tptp/Problems/COM/COM022+1.p', m__)).
fof(f249, plain, ! [X0, X1] : (~ sP4(X0, X1) | aElement0(sK20(X0, X1)) | ~ aElement0(sK19(X0, X1))), inference(subsumption_resolution, [], [f248, f144])).
fof(f144, plain, aRewritingSystem0(xR), inference(cnf_transformation, [], [f15])).
fof(f15, plain, aRewritingSystem0(xR), file('/home/ubuntu/library/tptp/Problems/COM/COM022+1.p', m__656)).
fof(f248, plain, ! [X0, X1] : (~ sP4(X0, X1) | aElement0(sK20(X0, X1)) | ~ aRewritingSystem0(xR) | ~ aElement0(sK19(X0, X1))), inference(resolution, [], [f156, f139])).
fof(f139, plain, ! [X2, X0, X1] : (~ aNormalFormOfIn0(X2, X0, X1) | aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ! [X0, X1] : (! [X2] : ((aNormalFormOfIn0(X2, X0, X1) | aReductOfIn0(sK16(X1, X2), X2, X1) | ~ sdtmndtasgtdt0(X0, X1, X2) | ~ aElement0(X2)) & ((! [X4] : ~ aReductOfIn0(X4, X2, X1) & sdtmndtasgtdt0(X0, X1, X2) & aElement0(X2)) | ~ aNormalFormOfIn0(X2, X0, X1))) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK16])], [f83, f84])).
fof(f84, plain, ! [X2, X1] : (? [X3] : aReductOfIn0(X3, X2, X1) => aReductOfIn0(sK16(X1, X2), X2, X1)), introduced(choice_axiom, [])).
fof(f83, plain, ! [X0, X1] : (! [X2] : ((aNormalFormOfIn0(X2, X0, X1) | ? [X3] : aReductOfIn0(X3, X2, X1) | ~ sdtmndtasgtdt0(X0, X1, X2) | ~ aElement0(X2)) & ((! [X4] : ~ aReductOfIn0(X4, X2, X1) & sdtmndtasgtdt0(X0, X1, X2) & aElement0(X2)) | ~ aNormalFormOfIn0(X2, X0, X1))) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(rectify, [], [f82])).
fof(f82, plain, ! [X0, X1] : (! [X2] : ((aNormalFormOfIn0(X2, X0, X1) | ? [X3] : aReductOfIn0(X3, X2, X1) | ~ sdtmndtasgtdt0(X0, X1, X2) | ~ aElement0(X2)) & ((! [X3] : ~ aReductOfIn0(X3, X2, X1) & sdtmndtasgtdt0(X0, X1, X2) & aElement0(X2)) | ~ aNormalFormOfIn0(X2, X0, X1))) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(flattening, [], [f81])).
fof(f81, plain, ! [X0, X1] : (! [X2] : ((aNormalFormOfIn0(X2, X0, X1) | (? [X3] : aReductOfIn0(X3, X2, X1) | ~ sdtmndtasgtdt0(X0, X1, X2) | ~ aElement0(X2))) & ((! [X3] : ~ aReductOfIn0(X3, X2, X1) & sdtmndtasgtdt0(X0, X1, X2) & aElement0(X2)) | ~ aNormalFormOfIn0(X2, X0, X1))) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ! [X0, X1] : (! [X2] : (aNormalFormOfIn0(X2, X0, X1) <=> (! [X3] : ~ aReductOfIn0(X3, X2, X1) & sdtmndtasgtdt0(X0, X1, X2) & aElement0(X2))) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(flattening, [], [f42])).
fof(f42, plain, ! [X0, X1] : (! [X2] : (aNormalFormOfIn0(X2, X0, X1) <=> (! [X3] : ~ aReductOfIn0(X3, X2, X1) & sdtmndtasgtdt0(X0, X1, X2) & aElement0(X2))) | (~ aRewritingSystem0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0, X1] : ((aRewritingSystem0(X1) & aElement0(X0)) => ! [X2] : (aNormalFormOfIn0(X2, X0, X1) <=> (~ ? [X3] : aReductOfIn0(X3, X2, X1) & sdtmndtasgtdt0(X0, X1, X2) & aElement0(X2)))), file('/home/ubuntu/library/tptp/Problems/COM/COM022+1.p', mNFRDef)).
fof(f156, plain, ! [X0, X1] : (aNormalFormOfIn0(sK20(X0, X1), sK19(X0, X1), xR) | ~ sP4(X0, X1)), inference(cnf_transformation, [], [f94])).
fof(f184, plain, (sP4(sK22, sK21) | ~ spl23_3), inference(avatar_component_clause, [], [f182])).
fof(f182, plain, (spl23_3 <=> sP4(sK22, sK21)), introduced(avatar_definition, [new_symbols(naming, [spl23_3])])).
fof(f367, plain, ~ spl23_12, inference(avatar_contradiction_clause, [], [f366])).
fof(f366, plain, ($false | ~ spl23_12), inference(subsumption_resolution, [], [f365, f144])).
fof(f365, plain, (~ aRewritingSystem0(xR) | ~ spl23_12), inference(subsumption_resolution, [], [f364, f148])).
fof(f148, plain, aElement0(xb), inference(cnf_transformation, [], [f17])).
fof(f17, plain, (aElement0(xc) & aElement0(xb) & aElement0(xa)), file('/home/ubuntu/library/tptp/Problems/COM/COM022+1.p', m__731)).
fof(f364, plain, (~ aElement0(xb) | ~ aRewritingSystem0(xR) | ~ spl23_12), inference(subsumption_resolution, [], [f363, f166])).
fof(f166, plain, sdtmndtasgtdt0(xa, xR, xb), inference(cnf_transformation, [], [f98])).
fof(f363, plain, (~ sdtmndtasgtdt0(xa, xR, xb) | ~ aElement0(xb) | ~ aRewritingSystem0(xR) | ~ spl23_12), inference(duplicate_literal_removal, [], [f362])).
fof(f362, plain, (~ sdtmndtasgtdt0(xa, xR, xb) | ~ aElement0(xb) | ~ aElement0(xb) | ~ aRewritingSystem0(xR) | ~ spl23_12), inference(resolution, [], [f348, f170])).
fof(f170, plain, ! [X2, X1] : (sdtmndtasgtdt0(X2, X1, X2) | ~ aElement0(X2) | ~ aRewritingSystem0(X1)), inference(duplicate_literal_removal, [], [f169])).
fof(f169, plain, ! [X2, X1] : (sdtmndtasgtdt0(X2, X1, X2) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X2)), inference(equality_resolution, [], [f107])).
fof(f107, plain, ! [X2, X0, X1] : (sdtmndtasgtdt0(X0, X1, X2) | ~ (X0 = X2) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f64])).
fof(f64, plain, ! [X0, X1, X2] : (((sdtmndtasgtdt0(X0, X1, X2) | (~ sdtmndtplgtdt0(X0, X1, X2) & ~ (X0 = X2))) & (sdtmndtplgtdt0(X0, X1, X2) | (X0 = X2) | ~ sdtmndtasgtdt0(X0, X1, X2))) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(flattening, [], [f63])).
fof(f63, plain, ! [X0, X1, X2] : (((sdtmndtasgtdt0(X0, X1, X2) | (~ sdtmndtplgtdt0(X0, X1, X2) & ~ (X0 = X2))) & ((sdtmndtplgtdt0(X0, X1, X2) | (X0 = X2)) | ~ sdtmndtasgtdt0(X0, X1, X2))) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(nnf_transformation, [], [f33])).
fof(f33, plain, ! [X0, X1, X2] : ((sdtmndtasgtdt0(X0, X1, X2) <=> (sdtmndtplgtdt0(X0, X1, X2) | (X0 = X2))) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(flattening, [], [f32])).
fof(f32, plain, ! [X0, X1, X2] : ((sdtmndtasgtdt0(X0, X1, X2) <=> (sdtmndtplgtdt0(X0, X1, X2) | (X0 = X2))) | (~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1, X2] : ((aElement0(X2) & aRewritingSystem0(X1) & aElement0(X0)) => (sdtmndtasgtdt0(X0, X1, X2) <=> (sdtmndtplgtdt0(X0, X1, X2) | (X0 = X2)))), file('/home/ubuntu/library/tptp/Problems/COM/COM022+1.p', mTCRDef)).
fof(f348, plain, (! [X0] : (~ sdtmndtasgtdt0(xb, xR, X0) | ~ sdtmndtasgtdt0(xa, xR, X0) | ~ aElement0(X0)) | ~ spl23_12), inference(backward_demodulation, [], [f168, f331])).
fof(f331, plain, ((xa = xc) | ~ spl23_12), inference(avatar_component_clause, [], [f329])).
fof(f329, plain, (spl23_12 <=> (xa = xc)), introduced(avatar_definition, [new_symbols(naming, [spl23_12])])).
fof(f332, plain, (spl23_2 | spl23_12), inference(avatar_split_clause, [], [f327, f329, f178])).
fof(f178, plain, (spl23_2 <=> sdtmndtplgtdt0(xa, xR, xc)), introduced(avatar_definition, [new_symbols(naming, [spl23_2])])).
fof(f327, plain, ((xa = xc) | sdtmndtplgtdt0(xa, xR, xc)), inference(subsumption_resolution, [], [f326, f147])).
fof(f147, plain, aElement0(xa), inference(cnf_transformation, [], [f17])).
fof(f326, plain, ((xa = xc) | sdtmndtplgtdt0(xa, xR, xc) | ~ aElement0(xa)), inference(subsumption_resolution, [], [f325, f144])).
fof(f325, plain, ((xa = xc) | sdtmndtplgtdt0(xa, xR, xc) | ~ aRewritingSystem0(xR) | ~ aElement0(xa)), inference(subsumption_resolution, [], [f302, f149])).
fof(f149, plain, aElement0(xc), inference(cnf_transformation, [], [f17])).
fof(f302, plain, ((xa = xc) | sdtmndtplgtdt0(xa, xR, xc) | ~ aElement0(xc) | ~ aRewritingSystem0(xR) | ~ aElement0(xa)), inference(resolution, [], [f106, f167])).
fof(f167, plain, sdtmndtasgtdt0(xa, xR, xc), inference(cnf_transformation, [], [f98])).
fof(f106, plain, ! [X2, X0, X1] : (~ sdtmndtasgtdt0(X0, X1, X2) | (X0 = X2) | sdtmndtplgtdt0(X0, X1, X2) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f64])).
fof(f324, plain, spl23_1, inference(avatar_contradiction_clause, [], [f323])).
fof(f323, plain, ($false | spl23_1), inference(subsumption_resolution, [], [f321, f167])).
fof(f321, plain, (~ sdtmndtasgtdt0(xa, xR, xc) | spl23_1), inference(backward_demodulation, [], [f243, f313])).
fof(f313, plain, ((xa = xb) | spl23_1), inference(subsumption_resolution, [], [f312, f147])).
fof(f312, plain, ((xa = xb) | ~ aElement0(xa) | spl23_1), inference(subsumption_resolution, [], [f311, f144])).
fof(f311, plain, ((xa = xb) | ~ aRewritingSystem0(xR) | ~ aElement0(xa) | spl23_1), inference(subsumption_resolution, [], [f310, f148])).
fof(f310, plain, ((xa = xb) | ~ aElement0(xb) | ~ aRewritingSystem0(xR) | ~ aElement0(xa) | spl23_1), inference(subsumption_resolution, [], [f301, f176])).
fof(f176, plain, (~ sdtmndtplgtdt0(xa, xR, xb) | spl23_1), inference(avatar_component_clause, [], [f174])).
fof(f174, plain, (spl23_1 <=> sdtmndtplgtdt0(xa, xR, xb)), introduced(avatar_definition, [new_symbols(naming, [spl23_1])])).
fof(f301, plain, ((xa = xb) | sdtmndtplgtdt0(xa, xR, xb) | ~ aElement0(xb) | ~ aRewritingSystem0(xR) | ~ aElement0(xa)), inference(resolution, [], [f106, f166])).
fof(f243, plain, ~ sdtmndtasgtdt0(xb, xR, xc), inference(subsumption_resolution, [], [f242, f144])).
fof(f242, plain, (~ sdtmndtasgtdt0(xb, xR, xc) | ~ aRewritingSystem0(xR)), inference(subsumption_resolution, [], [f238, f149])).
fof(f238, plain, (~ sdtmndtasgtdt0(xb, xR, xc) | ~ aElement0(xc) | ~ aRewritingSystem0(xR)), inference(duplicate_literal_removal, [], [f237])).
fof(f237, plain, (~ sdtmndtasgtdt0(xb, xR, xc) | ~ aElement0(xc) | ~ aElement0(xc) | ~ aRewritingSystem0(xR)), inference(resolution, [], [f168, f170])).
fof(f185, plain, (~ spl23_1 | ~ spl23_2 | spl23_3), inference(avatar_split_clause, [], [f165, f182, f178, f174])).
fof(f165, plain, (sP4(sK22, sK21) | ~ sdtmndtplgtdt0(xa, xR, xc) | ~ sdtmndtplgtdt0(xa, xR, xb)), inference(cnf_transformation, [], [f98])).