fof(f6024, plain, $false, inference(avatar_sat_refutation, [], [f416, f3618, f3643, f3797, f4495, f4535, f4538, f5493, f5500, f5916, f5932, f6013])).
fof(f6013, plain, (~ spl35_33 | ~ spl35_124 | ~ spl35_154), inference(avatar_contradiction_clause, [], [f6012])).
fof(f6012, plain, ($false | (~ spl35_33 | ~ spl35_124 | ~ spl35_154)), inference(subsumption_resolution, [], [f5949, f5701])).
fof(f5701, plain, (sP6(xb, xw) | (~ spl35_124 | ~ spl35_154)), inference(subsumption_resolution, [], [f5699, f3630])).
fof(f3630, plain, (aElement0(xw) | ~ spl35_124), inference(avatar_component_clause, [], [f3629])).
fof(f3629, plain, (spl35_124 <=> aElement0(xw)), introduced(avatar_definition, [new_symbols(naming, [spl35_124])])).
fof(f5699, plain, (~ aElement0(xw) | sP6(xb, xw) | ~ spl35_154), inference(resolution, [], [f5502, f4524])).
fof(f4524, plain, (! [X8] : (sP7(X8, xu) | ~ aElement0(X8) | sP6(xb, X8)) | ~ spl35_154), inference(avatar_component_clause, [], [f4523])).
fof(f4523, plain, (spl35_154 <=> ! [X8] : (sP6(xb, X8) | ~ aElement0(X8) | sP7(X8, xu))), introduced(avatar_definition, [new_symbols(naming, [spl35_154])])).
fof(f5502, plain, ~ sP7(xw, xu), inference(resolution, [], [f253, f208])).
fof(f208, plain, ! [X0, X1] : (~ sdtmndtasgtdt0(X1, xR, X0) | ~ sP7(X0, X1)), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ! [X0, X1] : ((~ sdtmndtasgtdt0(X1, xR, X0) & ~ sdtmndtplgtdt0(X1, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, X1, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, X1, xR) & ~ (X0 = X1)) | ~ sP7(X0, X1)), inference(rectify, [], [f113])).
fof(f113, plain, ! [X1, X0] : ((~ sdtmndtasgtdt0(X0, xR, X1) & ~ sdtmndtplgtdt0(X0, xR, X1) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X1) | ~ aReductOfIn0(X4, X0, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X1, X0, xR) & ~ (X0 = X1)) | ~ sP7(X1, X0)), inference(nnf_transformation, [], [f72])).
fof(f72, plain, ! [X1, X0] : ((~ sdtmndtasgtdt0(X0, xR, X1) & ~ sdtmndtplgtdt0(X0, xR, X1) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X1) | ~ aReductOfIn0(X4, X0, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X1, X0, xR) & ~ (X0 = X1)) | ~ sP7(X1, X0)), inference(usedef, [], [e72])).
fof(e72, plain, ! [X1, X0] : (sP7(X1, X0) <=> (~ sdtmndtasgtdt0(X0, xR, X1) & ~ sdtmndtplgtdt0(X0, xR, X1) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X1) | ~ aReductOfIn0(X4, X0, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X1, X0, xR) & ~ (X0 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f253, plain, sdtmndtasgtdt0(xu, xR, xw), inference(cnf_transformation, [], [f133])).
fof(f133, plain, (sdtmndtasgtdt0(xv, xR, xw) & ((sdtmndtplgtdt0(xv, xR, xw) & ((sdtmndtplgtdt0(sK32, xR, xw) & aReductOfIn0(sK32, xv, xR) & aElement0(sK32)) | aReductOfIn0(xw, xv, xR))) | (xv = xw)) & sdtmndtasgtdt0(xu, xR, xw) & ((sdtmndtplgtdt0(xu, xR, xw) & ((sdtmndtplgtdt0(sK33, xR, xw) & aReductOfIn0(sK33, xu, xR) & aElement0(sK33)) | aReductOfIn0(xw, xu, xR))) | (xu = xw)) & aElement0(xw)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK32, sK33])], [f33, f132, f131])).
fof(f131, plain, (? [X0] : (sdtmndtplgtdt0(X0, xR, xw) & aReductOfIn0(X0, xv, xR) & aElement0(X0)) => (sdtmndtplgtdt0(sK32, xR, xw) & aReductOfIn0(sK32, xv, xR) & aElement0(sK32))), introduced(choice_axiom, [])).
fof(f132, plain, (? [X1] : (sdtmndtplgtdt0(X1, xR, xw) & aReductOfIn0(X1, xu, xR) & aElement0(X1)) => (sdtmndtplgtdt0(sK33, xR, xw) & aReductOfIn0(sK33, xu, xR) & aElement0(sK33))), introduced(choice_axiom, [])).
fof(f33, plain, (sdtmndtasgtdt0(xv, xR, xw) & ((sdtmndtplgtdt0(xv, xR, xw) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xw) & aReductOfIn0(X0, xv, xR) & aElement0(X0)) | aReductOfIn0(xw, xv, xR))) | (xv = xw)) & sdtmndtasgtdt0(xu, xR, xw) & ((sdtmndtplgtdt0(xu, xR, xw) & (? [X1] : (sdtmndtplgtdt0(X1, xR, xw) & aReductOfIn0(X1, xu, xR) & aElement0(X1)) | aReductOfIn0(xw, xu, xR))) | (xu = xw)) & aElement0(xw)), inference(rectify, [], [f22])).
fof(f22, plain, (sdtmndtasgtdt0(xv, xR, xw) & ((sdtmndtplgtdt0(xv, xR, xw) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xw) & aReductOfIn0(X0, xv, xR) & aElement0(X0)) | aReductOfIn0(xw, xv, xR))) | (xv = xw)) & sdtmndtasgtdt0(xu, xR, xw) & ((sdtmndtplgtdt0(xu, xR, xw) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xw) & aReductOfIn0(X0, xu, xR) & aElement0(X0)) | aReductOfIn0(xw, xu, xR))) | (xu = xw)) & aElement0(xw)), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', m__799)).
fof(f5949, plain, (~ sP6(xb, xw) | ~ spl35_33), inference(backward_demodulation, [], [f664, f439])).
fof(f439, plain, ((xw = xd) | ~ spl35_33), inference(avatar_component_clause, [], [f437])).
fof(f437, plain, (spl35_33 <=> (xw = xd)), introduced(avatar_definition, [new_symbols(naming, [spl35_33])])).
fof(f664, plain, ~ sP6(xb, xd), inference(duplicate_literal_removal, [], [f663])).
fof(f663, plain, (~ sP6(xb, xd) | ~ sP6(xb, xd)), inference(resolution, [], [f640, f639])).
fof(f639, plain, ! [X3] : (sP8(sK25(X3, xd)) | ~ sP6(X3, xd)), inference(subsumption_resolution, [], [f638, f209])).
fof(f209, plain, ! [X0, X1] : (aElement0(sK25(X0, X1)) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ! [X0, X1] : ((sdtmndtasgtdt0(X0, xR, sK25(X0, X1)) & ((sdtmndtplgtdt0(X0, xR, sK25(X0, X1)) & ((sdtmndtplgtdt0(sK26(X0, X1), xR, sK25(X0, X1)) & aReductOfIn0(sK26(X0, X1), X0, xR) & aElement0(sK26(X0, X1))) | aReductOfIn0(sK25(X0, X1), X0, xR))) | (sK25(X0, X1) = X0)) & sdtmndtasgtdt0(X1, xR, sK25(X0, X1)) & sP5(sK25(X0, X1), X1) & aElement0(sK25(X0, X1))) | ~ sP6(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK25, sK26])], [f116, f118, f117])).
fof(f117, plain, ! [X1, X0] : (? [X2] : (sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X1, xR, X2) & sP5(X2, X1) & aElement0(X2)) => (sdtmndtasgtdt0(X0, xR, sK25(X0, X1)) & ((sdtmndtplgtdt0(X0, xR, sK25(X0, X1)) & (? [X3] : (sdtmndtplgtdt0(X3, xR, sK25(X0, X1)) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(sK25(X0, X1), X0, xR))) | (sK25(X0, X1) = X0)) & sdtmndtasgtdt0(X1, xR, sK25(X0, X1)) & sP5(sK25(X0, X1), X1) & aElement0(sK25(X0, X1)))), introduced(choice_axiom, [])).
fof(f118, plain, ! [X1, X0] : (? [X3] : (sdtmndtplgtdt0(X3, xR, sK25(X0, X1)) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) => (sdtmndtplgtdt0(sK26(X0, X1), xR, sK25(X0, X1)) & aReductOfIn0(sK26(X0, X1), X0, xR) & aElement0(sK26(X0, X1)))), introduced(choice_axiom, [])).
fof(f116, plain, ! [X0, X1] : (? [X2] : (sdtmndtasgtdt0(X0, xR, X2) & ((sdtmndtplgtdt0(X0, xR, X2) & (? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR))) | (X0 = X2)) & sdtmndtasgtdt0(X1, xR, X2) & sP5(X2, X1) & aElement0(X2)) | ~ sP6(X0, X1)), inference(rectify, [], [f115])).
fof(f115, plain, ! [X2, X1] : (? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & sP5(X5, X1) & aElement0(X5)) | ~ sP6(X2, X1)), inference(nnf_transformation, [], [f71])).
fof(f71, plain, ! [X2, X1] : (? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & sP5(X5, X1) & aElement0(X5)) | ~ sP6(X2, X1)), inference(usedef, [], [e71])).
fof(e71, plain, ! [X2, X1] : (sP6(X2, X1) <=> ? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & sP5(X5, X1) & aElement0(X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f638, plain, ! [X3] : (~ sP6(X3, xd) | sP8(sK25(X3, xd)) | ~ aElement0(sK25(X3, xd))), inference(resolution, [], [f211, f276])).
fof(f276, plain, ! [X0] : (~ sdtmndtasgtdt0(xd, xR, X0) | sP8(X0) | ~ aElement0(X0)), inference(cnf_transformation, [], [f75])).
fof(f75, plain, ! [X0] : ((~ sdtmndtasgtdt0(xd, xR, X0) & ~ sdtmndtplgtdt0(xd, xR, X0) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, X0) | ~ aReductOfIn0(X1, xd, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(X0, xd, xR) & ~ (xd = X0)) | sP8(X0) | ~ aElement0(X0)), inference(definition_folding, [], [f61, e74])).
fof(f74, plain, ! [X0] : ((~ sdtmndtasgtdt0(xb, xR, X0) & ~ sdtmndtplgtdt0(xb, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, xb, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, xb, xR) & ~ (xb = X0)) | ~ sP8(X0)), inference(usedef, [], [e74])).
fof(e74, plain, ! [X0] : (sP8(X0) <=> (~ sdtmndtasgtdt0(xb, xR, X0) & ~ sdtmndtplgtdt0(xb, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, xb, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, xb, xR) & ~ (xb = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f61, plain, ! [X0] : ((~ sdtmndtasgtdt0(xd, xR, X0) & ~ sdtmndtplgtdt0(xd, xR, X0) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, X0) | ~ aReductOfIn0(X1, xd, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(X0, xd, xR) & ~ (xd = X0)) | (~ sdtmndtasgtdt0(xb, xR, X0) & ~ sdtmndtplgtdt0(xb, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, xb, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, xb, xR) & ~ (xb = X0)) | ~ aElement0(X0)), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ~ ? [X0] : ((sdtmndtasgtdt0(xd, xR, X0) | sdtmndtplgtdt0(xd, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xd, xR) & aElement0(X1)) | aReductOfIn0(X0, xd, xR) | (xd = X0)) & (sdtmndtasgtdt0(xb, xR, X0) | sdtmndtplgtdt0(xb, xR, X0) | ? [X2] : (sdtmndtplgtdt0(X2, xR, X0) & aReductOfIn0(X2, xb, xR) & aElement0(X2)) | aReductOfIn0(X0, xb, xR) | (xb = X0)) & aElement0(X0)), inference(rectify, [], [f25])).
fof(f25, plain, ~ ? [X0] : ((sdtmndtasgtdt0(xd, xR, X0) | sdtmndtplgtdt0(xd, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xd, xR) & aElement0(X1)) | aReductOfIn0(X0, xd, xR) | (xd = X0)) & (sdtmndtasgtdt0(xb, xR, X0) | sdtmndtplgtdt0(xb, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xb, xR) & aElement0(X1)) | aReductOfIn0(X0, xb, xR) | (xb = X0)) & aElement0(X0)), inference(negated_conjecture, [], [f24])).
fof(f24, plain, ~ ? [X0] : ((sdtmndtasgtdt0(xd, xR, X0) | sdtmndtplgtdt0(xd, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xd, xR) & aElement0(X1)) | aReductOfIn0(X0, xd, xR) | (xd = X0)) & (sdtmndtasgtdt0(xb, xR, X0) | sdtmndtplgtdt0(xb, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xb, xR) & aElement0(X1)) | aReductOfIn0(X0, xb, xR) | (xb = X0)) & aElement0(X0)), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', m__)).
fof(f211, plain, ! [X0, X1] : (sdtmndtasgtdt0(X1, xR, sK25(X0, X1)) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f119])).
fof(f640, plain, ! [X0] : (~ sP8(sK25(xb, X0)) | ~ sP6(xb, X0)), inference(resolution, [], [f216, f271])).
fof(f271, plain, ! [X0] : (~ sdtmndtasgtdt0(xb, xR, X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ! [X0] : ((~ sdtmndtasgtdt0(xb, xR, X0) & ~ sdtmndtplgtdt0(xb, xR, X0) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, X0) | ~ aReductOfIn0(X1, xb, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(X0, xb, xR) & ~ (xb = X0)) | ~ sP8(X0)), inference(rectify, [], [f136])).
fof(f136, plain, ! [X0] : ((~ sdtmndtasgtdt0(xb, xR, X0) & ~ sdtmndtplgtdt0(xb, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, xb, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, xb, xR) & ~ (xb = X0)) | ~ sP8(X0)), inference(nnf_transformation, [], [f74])).
fof(f216, plain, ! [X0, X1] : (sdtmndtasgtdt0(X0, xR, sK25(X0, X1)) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f119])).
fof(f5932, plain, (spl35_34 | spl35_33 | ~ spl35_124), inference(avatar_split_clause, [], [f5931, f3629, f437, f441])).
fof(f441, plain, (spl35_34 <=> sdtmndtplgtdt0(xw, xR, xd)), introduced(avatar_definition, [new_symbols(naming, [spl35_34])])).
fof(f5931, plain, ((xw = xd) | sdtmndtplgtdt0(xw, xR, xd) | ~ spl35_124), inference(subsumption_resolution, [], [f5930, f3630])).
fof(f5930, plain, ((xw = xd) | sdtmndtplgtdt0(xw, xR, xd) | ~ aElement0(xw)), inference(subsumption_resolution, [], [f5929, f183])).
fof(f183, plain, aRewritingSystem0(xR), inference(cnf_transformation, [], [f15])).
fof(f15, plain, aRewritingSystem0(xR), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', m__656)).
fof(f5929, plain, ((xw = xd) | sdtmndtplgtdt0(xw, xR, xd) | ~ aRewritingSystem0(xR) | ~ aElement0(xw)), inference(subsumption_resolution, [], [f1040, f259])).
fof(f259, plain, aElement0(xd), inference(cnf_transformation, [], [f135])).
fof(f135, plain, (aNormalFormOfIn0(xd, xw, xR) & ! [X0] : ~ aReductOfIn0(X0, xd, xR) & sdtmndtasgtdt0(xw, xR, xd) & ((sdtmndtplgtdt0(xw, xR, xd) & ((sdtmndtplgtdt0(sK34, xR, xd) & aReductOfIn0(sK34, xw, xR) & aElement0(sK34)) | aReductOfIn0(xd, xw, xR))) | (xw = xd)) & aElement0(xd)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK34])], [f60, f134])).
fof(f134, plain, (? [X1] : (sdtmndtplgtdt0(X1, xR, xd) & aReductOfIn0(X1, xw, xR) & aElement0(X1)) => (sdtmndtplgtdt0(sK34, xR, xd) & aReductOfIn0(sK34, xw, xR) & aElement0(sK34))), introduced(choice_axiom, [])).
fof(f60, plain, (aNormalFormOfIn0(xd, xw, xR) & ! [X0] : ~ aReductOfIn0(X0, xd, xR) & sdtmndtasgtdt0(xw, xR, xd) & ((sdtmndtplgtdt0(xw, xR, xd) & (? [X1] : (sdtmndtplgtdt0(X1, xR, xd) & aReductOfIn0(X1, xw, xR) & aElement0(X1)) | aReductOfIn0(xd, xw, xR))) | (xw = xd)) & aElement0(xd)), inference(ennf_transformation, [], [f34])).
fof(f34, plain, (aNormalFormOfIn0(xd, xw, xR) & ~ ? [X0] : aReductOfIn0(X0, xd, xR) & sdtmndtasgtdt0(xw, xR, xd) & ((sdtmndtplgtdt0(xw, xR, xd) & (? [X1] : (sdtmndtplgtdt0(X1, xR, xd) & aReductOfIn0(X1, xw, xR) & aElement0(X1)) | aReductOfIn0(xd, xw, xR))) | (xw = xd)) & aElement0(xd)), inference(rectify, [], [f23])).
fof(f23, plain, (aNormalFormOfIn0(xd, xw, xR) & ~ ? [X0] : aReductOfIn0(X0, xd, xR) & sdtmndtasgtdt0(xw, xR, xd) & ((sdtmndtplgtdt0(xw, xR, xd) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xd) & aReductOfIn0(X0, xw, xR) & aElement0(X0)) | aReductOfIn0(xd, xw, xR))) | (xw = xd)) & aElement0(xd)), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', m__818)).
fof(f1040, plain, ((xw = xd) | sdtmndtplgtdt0(xw, xR, xd) | ~ aElement0(xd) | ~ aRewritingSystem0(xR) | ~ aElement0(xw)), inference(resolution, [], [f145, f264])).
fof(f264, plain, sdtmndtasgtdt0(xw, xR, xd), inference(cnf_transformation, [], [f135])).
fof(f145, plain, ! [X2, X0, X1] : (~ sdtmndtasgtdt0(X0, X1, X2) | (X0 = X2) | sdtmndtplgtdt0(X0, X1, X2) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f82])).
fof(f82, plain, ! [X0, X1, X2] : (((sdtmndtasgtdt0(X0, X1, X2) | (~ sdtmndtplgtdt0(X0, X1, X2) & ~ (X0 = X2))) & (sdtmndtplgtdt0(X0, X1, X2) | (X0 = X2) | ~ sdtmndtasgtdt0(X0, X1, X2))) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(flattening, [], [f81])).
fof(f81, plain, ! [X0, X1, X2] : (((sdtmndtasgtdt0(X0, X1, X2) | (~ sdtmndtplgtdt0(X0, X1, X2) & ~ (X0 = X2))) & ((sdtmndtplgtdt0(X0, X1, X2) | (X0 = X2)) | ~ sdtmndtasgtdt0(X0, X1, X2))) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(nnf_transformation, [], [f43])).
fof(f43, plain, ! [X0, X1, X2] : ((sdtmndtasgtdt0(X0, X1, X2) <=> (sdtmndtplgtdt0(X0, X1, X2) | (X0 = X2))) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(flattening, [], [f42])).
fof(f42, plain, ! [X0, X1, X2] : ((sdtmndtasgtdt0(X0, X1, X2) <=> (sdtmndtplgtdt0(X0, X1, X2) | (X0 = X2))) | (~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1, X2] : ((aElement0(X2) & aRewritingSystem0(X1) & aElement0(X0)) => (sdtmndtasgtdt0(X0, X1, X2) <=> (sdtmndtplgtdt0(X0, X1, X2) | (X0 = X2)))), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', mTCRDef)).
fof(f5916, plain, (~ spl35_28 | ~ spl35_34 | ~ spl35_128 | ~ spl35_154), inference(avatar_contradiction_clause, [], [f5915])).
fof(f5915, plain, ($false | (~ spl35_28 | ~ spl35_34 | ~ spl35_128 | ~ spl35_154)), inference(subsumption_resolution, [], [f5914, f664])).
fof(f5914, plain, (sP6(xb, xd) | (~ spl35_28 | ~ spl35_34 | ~ spl35_128 | ~ spl35_154)), inference(subsumption_resolution, [], [f5912, f259])).
fof(f5912, plain, (~ aElement0(xd) | sP6(xb, xd) | (~ spl35_28 | ~ spl35_34 | ~ spl35_128 | ~ spl35_154)), inference(resolution, [], [f5881, f4524])).
fof(f5881, plain, (~ sP7(xd, xu) | (~ spl35_28 | ~ spl35_34 | ~ spl35_128)), inference(resolution, [], [f5844, f207])).
fof(f207, plain, ! [X0, X1] : (~ sdtmndtplgtdt0(X1, xR, X0) | ~ sP7(X0, X1)), inference(cnf_transformation, [], [f114])).
fof(f5844, plain, (sdtmndtplgtdt0(xu, xR, xd) | (~ spl35_28 | ~ spl35_34 | ~ spl35_128)), inference(subsumption_resolution, [], [f5838, f3792])).
fof(f3792, plain, (aElement0(xu) | ~ spl35_128), inference(avatar_component_clause, [], [f3791])).
fof(f3791, plain, (spl35_128 <=> aElement0(xu)), introduced(avatar_definition, [new_symbols(naming, [spl35_128])])).
fof(f5838, plain, (sdtmndtplgtdt0(xu, xR, xd) | ~ aElement0(xu) | (~ spl35_28 | ~ spl35_34)), inference(resolution, [], [f1502, f415])).
fof(f415, plain, (sdtmndtplgtdt0(xu, xR, xw) | ~ spl35_28), inference(avatar_component_clause, [], [f413])).
fof(f413, plain, (spl35_28 <=> sdtmndtplgtdt0(xu, xR, xw)), introduced(avatar_definition, [new_symbols(naming, [spl35_28])])).
fof(f1502, plain, (! [X18] : (~ sdtmndtplgtdt0(X18, xR, xw) | sdtmndtplgtdt0(X18, xR, xd) | ~ aElement0(X18)) | ~ spl35_34), inference(subsumption_resolution, [], [f1501, f183])).
fof(f1501, plain, (! [X18] : (sdtmndtplgtdt0(X18, xR, xd) | ~ sdtmndtplgtdt0(X18, xR, xw) | ~ aRewritingSystem0(xR) | ~ aElement0(X18)) | ~ spl35_34), inference(subsumption_resolution, [], [f1500, f248])).
fof(f248, plain, aElement0(xw), inference(cnf_transformation, [], [f133])).
fof(f1500, plain, (! [X18] : (sdtmndtplgtdt0(X18, xR, xd) | ~ sdtmndtplgtdt0(X18, xR, xw) | ~ aElement0(xw) | ~ aRewritingSystem0(xR) | ~ aElement0(X18)) | ~ spl35_34), inference(subsumption_resolution, [], [f1463, f259])).
fof(f1463, plain, (! [X18] : (sdtmndtplgtdt0(X18, xR, xd) | ~ sdtmndtplgtdt0(X18, xR, xw) | ~ aElement0(xd) | ~ aElement0(xw) | ~ aRewritingSystem0(xR) | ~ aElement0(X18)) | ~ spl35_34), inference(resolution, [], [f144, f443])).
fof(f443, plain, (sdtmndtplgtdt0(xw, xR, xd) | ~ spl35_34), inference(avatar_component_clause, [], [f441])).
fof(f144, plain, ! [X2, X0, X3, X1] : (~ sdtmndtplgtdt0(X2, X1, X3) | sdtmndtplgtdt0(X0, X1, X3) | ~ sdtmndtplgtdt0(X0, X1, X2) | ~ aElement0(X3) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f41])).
fof(f41, plain, ! [X0, X1, X2, X3] : (sdtmndtplgtdt0(X0, X1, X3) | ~ sdtmndtplgtdt0(X2, X1, X3) | ~ sdtmndtplgtdt0(X0, X1, X2) | ~ aElement0(X3) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(flattening, [], [f40])).
fof(f40, plain, ! [X0, X1, X2, X3] : ((sdtmndtplgtdt0(X0, X1, X3) | (~ sdtmndtplgtdt0(X2, X1, X3) | ~ sdtmndtplgtdt0(X0, X1, X2))) | (~ aElement0(X3) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ! [X0, X1, X2, X3] : ((aElement0(X3) & aElement0(X2) & aRewritingSystem0(X1) & aElement0(X0)) => ((sdtmndtplgtdt0(X2, X1, X3) & sdtmndtplgtdt0(X0, X1, X2)) => sdtmndtplgtdt0(X0, X1, X3))), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', mTCTrans)).
fof(f5500, plain, (spl35_22 | spl35_21 | ~ spl35_124), inference(avatar_split_clause, [], [f5499, f3629, f381, f385])).
fof(f385, plain, (spl35_22 <=> sdtmndtplgtdt0(xv, xR, xw)), introduced(avatar_definition, [new_symbols(naming, [spl35_22])])).
fof(f381, plain, (spl35_21 <=> (xv = xw)), introduced(avatar_definition, [new_symbols(naming, [spl35_21])])).
fof(f5499, plain, ((xv = xw) | sdtmndtplgtdt0(xv, xR, xw) | ~ spl35_124), inference(subsumption_resolution, [], [f4514, f3630])).
fof(f4514, plain, ((xv = xw) | sdtmndtplgtdt0(xv, xR, xw) | ~ aElement0(xw)), inference(subsumption_resolution, [], [f4513, f241])).
fof(f241, plain, aElement0(xv), inference(cnf_transformation, [], [f130])).
fof(f130, plain, (sdtmndtasgtdt0(xv, xR, xc) & ((sdtmndtplgtdt0(xv, xR, xc) & ((sdtmndtplgtdt0(sK31, xR, xc) & aReductOfIn0(sK31, xv, xR) & aElement0(sK31)) | aReductOfIn0(xc, xv, xR))) | (xc = xv)) & aReductOfIn0(xv, xa, xR) & aElement0(xv)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK31])], [f21, f129])).
fof(f129, plain, (? [X0] : (sdtmndtplgtdt0(X0, xR, xc) & aReductOfIn0(X0, xv, xR) & aElement0(X0)) => (sdtmndtplgtdt0(sK31, xR, xc) & aReductOfIn0(sK31, xv, xR) & aElement0(sK31))), introduced(choice_axiom, [])).
fof(f21, plain, (sdtmndtasgtdt0(xv, xR, xc) & ((sdtmndtplgtdt0(xv, xR, xc) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xc) & aReductOfIn0(X0, xv, xR) & aElement0(X0)) | aReductOfIn0(xc, xv, xR))) | (xc = xv)) & aReductOfIn0(xv, xa, xR) & aElement0(xv)), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', m__779)).
fof(f4513, plain, ((xv = xw) | sdtmndtplgtdt0(xv, xR, xw) | ~ aElement0(xw) | ~ aElement0(xv)), inference(subsumption_resolution, [], [f1039, f183])).
fof(f1039, plain, ((xv = xw) | sdtmndtplgtdt0(xv, xR, xw) | ~ aElement0(xw) | ~ aRewritingSystem0(xR) | ~ aElement0(xv)), inference(resolution, [], [f145, f258])).
fof(f258, plain, sdtmndtasgtdt0(xv, xR, xw), inference(cnf_transformation, [], [f133])).
fof(f5493, plain, (~ spl35_21 | ~ spl35_128 | ~ spl35_154), inference(avatar_contradiction_clause, [], [f5492])).
fof(f5492, plain, ($false | (~ spl35_21 | ~ spl35_128 | ~ spl35_154)), inference(subsumption_resolution, [], [f5491, f664])).
fof(f5491, plain, (sP6(xb, xd) | (~ spl35_21 | ~ spl35_128 | ~ spl35_154)), inference(subsumption_resolution, [], [f5488, f259])).
fof(f5488, plain, (~ aElement0(xd) | sP6(xb, xd) | (~ spl35_21 | ~ spl35_128 | ~ spl35_154)), inference(resolution, [], [f5460, f4524])).
fof(f5460, plain, (~ sP7(xd, xu) | (~ spl35_21 | ~ spl35_128)), inference(resolution, [], [f5396, f208])).
fof(f5396, plain, (sdtmndtasgtdt0(xu, xR, xd) | (~ spl35_21 | ~ spl35_128)), inference(subsumption_resolution, [], [f5389, f3792])).
fof(f5389, plain, (sdtmndtasgtdt0(xu, xR, xd) | ~ aElement0(xu) | ~ spl35_21), inference(resolution, [], [f4673, f4576])).
fof(f4576, plain, (sdtmndtasgtdt0(xu, xR, xv) | ~ spl35_21), inference(forward_demodulation, [], [f253, f383])).
fof(f383, plain, ((xv = xw) | ~ spl35_21), inference(avatar_component_clause, [], [f381])).
fof(f4673, plain, (! [X12] : (~ sdtmndtasgtdt0(X12, xR, xv) | sdtmndtasgtdt0(X12, xR, xd) | ~ aElement0(X12)) | ~ spl35_21), inference(forward_demodulation, [], [f1554, f383])).
fof(f1554, plain, ! [X12] : (sdtmndtasgtdt0(X12, xR, xd) | ~ sdtmndtasgtdt0(X12, xR, xw) | ~ aElement0(X12)), inference(subsumption_resolution, [], [f1553, f183])).
fof(f1553, plain, ! [X12] : (sdtmndtasgtdt0(X12, xR, xd) | ~ sdtmndtasgtdt0(X12, xR, xw) | ~ aRewritingSystem0(xR) | ~ aElement0(X12)), inference(subsumption_resolution, [], [f1552, f248])).
fof(f1552, plain, ! [X12] : (sdtmndtasgtdt0(X12, xR, xd) | ~ sdtmndtasgtdt0(X12, xR, xw) | ~ aElement0(xw) | ~ aRewritingSystem0(xR) | ~ aElement0(X12)), inference(subsumption_resolution, [], [f1526, f259])).
fof(f1526, plain, ! [X12] : (sdtmndtasgtdt0(X12, xR, xd) | ~ sdtmndtasgtdt0(X12, xR, xw) | ~ aElement0(xd) | ~ aElement0(xw) | ~ aRewritingSystem0(xR) | ~ aElement0(X12)), inference(resolution, [], [f148, f264])).
fof(f148, plain, ! [X2, X0, X3, X1] : (~ sdtmndtasgtdt0(X2, X1, X3) | sdtmndtasgtdt0(X0, X1, X3) | ~ sdtmndtasgtdt0(X0, X1, X2) | ~ aElement0(X3) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f45])).
fof(f45, plain, ! [X0, X1, X2, X3] : (sdtmndtasgtdt0(X0, X1, X3) | ~ sdtmndtasgtdt0(X2, X1, X3) | ~ sdtmndtasgtdt0(X0, X1, X2) | ~ aElement0(X3) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0)), inference(flattening, [], [f44])).
fof(f44, plain, ! [X0, X1, X2, X3] : ((sdtmndtasgtdt0(X0, X1, X3) | (~ sdtmndtasgtdt0(X2, X1, X3) | ~ sdtmndtasgtdt0(X0, X1, X2))) | (~ aElement0(X3) | ~ aElement0(X2) | ~ aRewritingSystem0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1, X2, X3] : ((aElement0(X3) & aElement0(X2) & aRewritingSystem0(X1) & aElement0(X0)) => ((sdtmndtasgtdt0(X2, X1, X3) & sdtmndtasgtdt0(X0, X1, X2)) => sdtmndtasgtdt0(X0, X1, X3))), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', mTCRTrans)).
fof(f4538, plain, (spl35_155 | ~ spl35_128), inference(avatar_split_clause, [], [f4537, f3791, f4526])).
fof(f4526, plain, (spl35_155 <=> iLess0(xu, xa)), introduced(avatar_definition, [new_symbols(naming, [spl35_155])])).
fof(f4537, plain, (iLess0(xu, xa) | ~ spl35_128), inference(subsumption_resolution, [], [f4536, f201])).
fof(f201, plain, aElement0(xa), inference(cnf_transformation, [], [f17])).
fof(f17, plain, (aElement0(xc) & aElement0(xb) & aElement0(xa)), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', m__731)).
fof(f4536, plain, (iLess0(xu, xa) | ~ aElement0(xa) | ~ spl35_128), inference(subsumption_resolution, [], [f3814, f3792])).
fof(f3814, plain, (iLess0(xu, xa) | ~ aElement0(xu) | ~ aElement0(xa)), inference(resolution, [], [f235, f197])).
fof(f197, plain, ! [X0, X1] : (~ aReductOfIn0(X1, X0, xR) | iLess0(X1, X0) | ~ aElement0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f112])).
fof(f112, plain, (isTerminating0(xR) & ! [X0, X1] : (iLess0(X1, X0) | (~ sdtmndtplgtdt0(X0, xR, X1) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X1) | ~ aReductOfIn0(X2, X0, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X1, X0, xR)) | ~ aElement0(X1) | ~ aElement0(X0)) & isLocallyConfluent0(xR) & ! [X3, X4, X5] : ((sdtmndtasgtdt0(X5, xR, sK23(X4, X5)) & ((sdtmndtplgtdt0(X5, xR, sK23(X4, X5)) & ((sdtmndtplgtdt0(sK24(X4, X5), xR, sK23(X4, X5)) & aReductOfIn0(sK24(X4, X5), X5, xR) & aElement0(sK24(X4, X5))) | aReductOfIn0(sK23(X4, X5), X5, xR))) | (sK23(X4, X5) = X5)) & sdtmndtasgtdt0(X4, xR, sK23(X4, X5)) & sP4(sK23(X4, X5), X4) & aElement0(sK23(X4, X5))) | ~ aReductOfIn0(X5, X3, xR) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK23, sK24])], [f69, f111, f110])).
fof(f110, plain, ! [X5, X4] : (? [X6] : (sdtmndtasgtdt0(X5, xR, X6) & ((sdtmndtplgtdt0(X5, xR, X6) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X6) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(X6, X5, xR))) | (X5 = X6)) & sdtmndtasgtdt0(X4, xR, X6) & sP4(X6, X4) & aElement0(X6)) => (sdtmndtasgtdt0(X5, xR, sK23(X4, X5)) & ((sdtmndtplgtdt0(X5, xR, sK23(X4, X5)) & (? [X7] : (sdtmndtplgtdt0(X7, xR, sK23(X4, X5)) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(sK23(X4, X5), X5, xR))) | (sK23(X4, X5) = X5)) & sdtmndtasgtdt0(X4, xR, sK23(X4, X5)) & sP4(sK23(X4, X5), X4) & aElement0(sK23(X4, X5)))), introduced(choice_axiom, [])).
fof(f111, plain, ! [X5, X4] : (? [X7] : (sdtmndtplgtdt0(X7, xR, sK23(X4, X5)) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) => (sdtmndtplgtdt0(sK24(X4, X5), xR, sK23(X4, X5)) & aReductOfIn0(sK24(X4, X5), X5, xR) & aElement0(sK24(X4, X5)))), introduced(choice_axiom, [])).
fof(f69, plain, (isTerminating0(xR) & ! [X0, X1] : (iLess0(X1, X0) | (~ sdtmndtplgtdt0(X0, xR, X1) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X1) | ~ aReductOfIn0(X2, X0, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X1, X0, xR)) | ~ aElement0(X1) | ~ aElement0(X0)) & isLocallyConfluent0(xR) & ! [X3, X4, X5] : (? [X6] : (sdtmndtasgtdt0(X5, xR, X6) & ((sdtmndtplgtdt0(X5, xR, X6) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X6) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(X6, X5, xR))) | (X5 = X6)) & sdtmndtasgtdt0(X4, xR, X6) & sP4(X6, X4) & aElement0(X6)) | ~ aReductOfIn0(X5, X3, xR) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3))), inference(definition_folding, [], [f57, e68])).
fof(f68, plain, ! [X6, X4] : ((sdtmndtplgtdt0(X4, xR, X6) & (? [X8] : (sdtmndtplgtdt0(X8, xR, X6) & aReductOfIn0(X8, X4, xR) & aElement0(X8)) | aReductOfIn0(X6, X4, xR))) | (X4 = X6) | ~ sP4(X6, X4)), inference(usedef, [], [e68])).
fof(e68, plain, ! [X6, X4] : (sP4(X6, X4) <=> ((sdtmndtplgtdt0(X4, xR, X6) & (? [X8] : (sdtmndtplgtdt0(X8, xR, X6) & aReductOfIn0(X8, X4, xR) & aElement0(X8)) | aReductOfIn0(X6, X4, xR))) | (X4 = X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f57, plain, (isTerminating0(xR) & ! [X0, X1] : (iLess0(X1, X0) | (~ sdtmndtplgtdt0(X0, xR, X1) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X1) | ~ aReductOfIn0(X2, X0, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X1, X0, xR)) | ~ aElement0(X1) | ~ aElement0(X0)) & isLocallyConfluent0(xR) & ! [X3, X4, X5] : (? [X6] : (sdtmndtasgtdt0(X5, xR, X6) & ((sdtmndtplgtdt0(X5, xR, X6) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X6) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(X6, X5, xR))) | (X5 = X6)) & sdtmndtasgtdt0(X4, xR, X6) & ((sdtmndtplgtdt0(X4, xR, X6) & (? [X8] : (sdtmndtplgtdt0(X8, xR, X6) & aReductOfIn0(X8, X4, xR) & aElement0(X8)) | aReductOfIn0(X6, X4, xR))) | (X4 = X6)) & aElement0(X6)) | ~ aReductOfIn0(X5, X3, xR) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3))), inference(flattening, [], [f56])).
fof(f56, plain, (isTerminating0(xR) & ! [X0, X1] : ((iLess0(X1, X0) | (~ sdtmndtplgtdt0(X0, xR, X1) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X1) | ~ aReductOfIn0(X2, X0, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X1, X0, xR))) | (~ aElement0(X1) | ~ aElement0(X0))) & isLocallyConfluent0(xR) & ! [X3, X4, X5] : (? [X6] : (sdtmndtasgtdt0(X5, xR, X6) & ((sdtmndtplgtdt0(X5, xR, X6) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X6) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(X6, X5, xR))) | (X5 = X6)) & sdtmndtasgtdt0(X4, xR, X6) & ((sdtmndtplgtdt0(X4, xR, X6) & (? [X8] : (sdtmndtplgtdt0(X8, xR, X6) & aReductOfIn0(X8, X4, xR) & aElement0(X8)) | aReductOfIn0(X6, X4, xR))) | (X4 = X6)) & aElement0(X6)) | (~ aReductOfIn0(X5, X3, xR) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3)))), inference(ennf_transformation, [], [f30])).
fof(f30, plain, (isTerminating0(xR) & ! [X0, X1] : ((aElement0(X1) & aElement0(X0)) => ((sdtmndtplgtdt0(X0, xR, X1) | ? [X2] : (sdtmndtplgtdt0(X2, xR, X1) & aReductOfIn0(X2, X0, xR) & aElement0(X2)) | aReductOfIn0(X1, X0, xR)) => iLess0(X1, X0))) & isLocallyConfluent0(xR) & ! [X3, X4, X5] : ((aReductOfIn0(X5, X3, xR) & aReductOfIn0(X4, X3, xR) & aElement0(X5) & aElement0(X4) & aElement0(X3)) => ? [X6] : (sdtmndtasgtdt0(X5, xR, X6) & ((sdtmndtplgtdt0(X5, xR, X6) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X6) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(X6, X5, xR))) | (X5 = X6)) & sdtmndtasgtdt0(X4, xR, X6) & ((sdtmndtplgtdt0(X4, xR, X6) & (? [X8] : (sdtmndtplgtdt0(X8, xR, X6) & aReductOfIn0(X8, X4, xR) & aElement0(X8)) | aReductOfIn0(X6, X4, xR))) | (X4 = X6)) & aElement0(X6)))), inference(rectify, [], [f16])).
fof(f16, plain, (isTerminating0(xR) & ! [X0, X1] : ((aElement0(X1) & aElement0(X0)) => ((sdtmndtplgtdt0(X0, xR, X1) | ? [X2] : (sdtmndtplgtdt0(X2, xR, X1) & aReductOfIn0(X2, X0, xR) & aElement0(X2)) | aReductOfIn0(X1, X0, xR)) => iLess0(X1, X0))) & isLocallyConfluent0(xR) & ! [X0, X1, X2] : ((aReductOfIn0(X2, X0, xR) & aReductOfIn0(X1, X0, xR) & aElement0(X2) & aElement0(X1) & aElement0(X0)) => ? [X3] : (sdtmndtasgtdt0(X2, xR, X3) & ((sdtmndtplgtdt0(X2, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X2, xR) & aElement0(X4)) | aReductOfIn0(X3, X2, xR))) | (X2 = X3)) & sdtmndtasgtdt0(X1, xR, X3) & ((sdtmndtplgtdt0(X1, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X1, xR) & aElement0(X4)) | aReductOfIn0(X3, X1, xR))) | (X1 = X3)) & aElement0(X3)))), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', m__656_01)).
fof(f235, plain, aReductOfIn0(xu, xa, xR), inference(cnf_transformation, [], [f128])).
fof(f128, plain, (sdtmndtasgtdt0(xu, xR, xb) & ((sdtmndtplgtdt0(xu, xR, xb) & ((sdtmndtplgtdt0(sK30, xR, xb) & aReductOfIn0(sK30, xu, xR) & aElement0(sK30)) | aReductOfIn0(xb, xu, xR))) | (xb = xu)) & aReductOfIn0(xu, xa, xR) & aElement0(xu)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK30])], [f20, f127])).
fof(f127, plain, (? [X0] : (sdtmndtplgtdt0(X0, xR, xb) & aReductOfIn0(X0, xu, xR) & aElement0(X0)) => (sdtmndtplgtdt0(sK30, xR, xb) & aReductOfIn0(sK30, xu, xR) & aElement0(sK30))), introduced(choice_axiom, [])).
fof(f20, plain, (sdtmndtasgtdt0(xu, xR, xb) & ((sdtmndtplgtdt0(xu, xR, xb) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xb) & aReductOfIn0(X0, xu, xR) & aElement0(X0)) | aReductOfIn0(xb, xu, xR))) | (xb = xu)) & aReductOfIn0(xu, xa, xR) & aElement0(xu)), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', m__755)).
fof(f4535, plain, (spl35_154 | ~ spl35_155 | ~ spl35_128), inference(avatar_split_clause, [], [f4534, f3791, f4526, f4523])).
fof(f4534, plain, (! [X0] : (~ iLess0(xu, xa) | sP6(xb, X0) | sP7(X0, xu) | ~ aElement0(X0)) | ~ spl35_128), inference(subsumption_resolution, [], [f4533, f3792])).
fof(f4533, plain, ! [X0] : (~ iLess0(xu, xa) | sP6(xb, X0) | sP7(X0, xu) | ~ aElement0(X0) | ~ aElement0(xu)), inference(subsumption_resolution, [], [f3849, f202])).
fof(f202, plain, aElement0(xb), inference(cnf_transformation, [], [f17])).
fof(f3849, plain, ! [X0] : (~ iLess0(xu, xa) | sP6(xb, X0) | sP7(X0, xu) | ~ aElement0(xb) | ~ aElement0(X0) | ~ aElement0(xu)), inference(resolution, [], [f240, f225])).
fof(f225, plain, ! [X2, X0, X1] : (~ sdtmndtasgtdt0(X0, xR, X2) | ~ iLess0(X0, xa) | sP6(X2, X1) | sP7(X1, X0) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f73])).
fof(f73, plain, ! [X0, X1, X2] : (sP6(X2, X1) | ~ iLess0(X0, xa) | (~ sdtmndtasgtdt0(X0, xR, X2) & ~ sdtmndtplgtdt0(X0, xR, X2) & ! [X3] : (~ sdtmndtplgtdt0(X3, xR, X2) | ~ aReductOfIn0(X3, X0, xR) | ~ aElement0(X3)) & ~ aReductOfIn0(X2, X0, xR) & ~ (X0 = X2)) | sP7(X1, X0) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0)), inference(definition_folding, [], [f59, e72, e71, e70])).
fof(f70, plain, ! [X5, X1] : ((sdtmndtplgtdt0(X1, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR))) | (X1 = X5) | ~ sP5(X5, X1)), inference(usedef, [], [e70])).
fof(e70, plain, ! [X5, X1] : (sP5(X5, X1) <=> ((sdtmndtplgtdt0(X1, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR))) | (X1 = X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f59, plain, ! [X0, X1, X2] : (? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & ((sdtmndtplgtdt0(X1, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR))) | (X1 = X5)) & aElement0(X5)) | ~ iLess0(X0, xa) | (~ sdtmndtasgtdt0(X0, xR, X2) & ~ sdtmndtplgtdt0(X0, xR, X2) & ! [X3] : (~ sdtmndtplgtdt0(X3, xR, X2) | ~ aReductOfIn0(X3, X0, xR) | ~ aElement0(X3)) & ~ aReductOfIn0(X2, X0, xR) & ~ (X0 = X2)) | (~ sdtmndtasgtdt0(X0, xR, X1) & ~ sdtmndtplgtdt0(X0, xR, X1) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X1) | ~ aReductOfIn0(X4, X0, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X1, X0, xR) & ~ (X0 = X1)) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0)), inference(flattening, [], [f58])).
fof(f58, plain, ! [X0, X1, X2] : ((? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & ((sdtmndtplgtdt0(X1, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR))) | (X1 = X5)) & aElement0(X5)) | ~ iLess0(X0, xa)) | ((~ sdtmndtasgtdt0(X0, xR, X2) & ~ sdtmndtplgtdt0(X0, xR, X2) & ! [X3] : (~ sdtmndtplgtdt0(X3, xR, X2) | ~ aReductOfIn0(X3, X0, xR) | ~ aElement0(X3)) & ~ aReductOfIn0(X2, X0, xR) & ~ (X0 = X2)) | (~ sdtmndtasgtdt0(X0, xR, X1) & ~ sdtmndtplgtdt0(X0, xR, X1) & ! [X4] : (~ sdtmndtplgtdt0(X4, xR, X1) | ~ aReductOfIn0(X4, X0, xR) | ~ aElement0(X4)) & ~ aReductOfIn0(X1, X0, xR) & ~ (X0 = X1)) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ! [X0, X1, X2] : (((sdtmndtasgtdt0(X0, xR, X2) | sdtmndtplgtdt0(X0, xR, X2) | ? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR) | (X0 = X2)) & (sdtmndtasgtdt0(X0, xR, X1) | sdtmndtplgtdt0(X0, xR, X1) | ? [X4] : (sdtmndtplgtdt0(X4, xR, X1) & aReductOfIn0(X4, X0, xR) & aElement0(X4)) | aReductOfIn0(X1, X0, xR) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0)) => (iLess0(X0, xa) => ? [X5] : (sdtmndtasgtdt0(X2, xR, X5) & ((sdtmndtplgtdt0(X2, xR, X5) & (? [X6] : (sdtmndtplgtdt0(X6, xR, X5) & aReductOfIn0(X6, X2, xR) & aElement0(X6)) | aReductOfIn0(X5, X2, xR))) | (X2 = X5)) & sdtmndtasgtdt0(X1, xR, X5) & ((sdtmndtplgtdt0(X1, xR, X5) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X5) & aReductOfIn0(X7, X1, xR) & aElement0(X7)) | aReductOfIn0(X5, X1, xR))) | (X1 = X5)) & aElement0(X5)))), inference(rectify, [], [f18])).
fof(f18, plain, ! [X0, X1, X2] : (((sdtmndtasgtdt0(X0, xR, X2) | sdtmndtplgtdt0(X0, xR, X2) | ? [X3] : (sdtmndtplgtdt0(X3, xR, X2) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X2, X0, xR) | (X0 = X2)) & (sdtmndtasgtdt0(X0, xR, X1) | sdtmndtplgtdt0(X0, xR, X1) | ? [X3] : (sdtmndtplgtdt0(X3, xR, X1) & aReductOfIn0(X3, X0, xR) & aElement0(X3)) | aReductOfIn0(X1, X0, xR) | (X0 = X1)) & aElement0(X2) & aElement0(X1) & aElement0(X0)) => (iLess0(X0, xa) => ? [X3] : (sdtmndtasgtdt0(X2, xR, X3) & ((sdtmndtplgtdt0(X2, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X2, xR) & aElement0(X4)) | aReductOfIn0(X3, X2, xR))) | (X2 = X3)) & sdtmndtasgtdt0(X1, xR, X3) & ((sdtmndtplgtdt0(X1, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X1, xR) & aElement0(X4)) | aReductOfIn0(X3, X1, xR))) | (X1 = X3)) & aElement0(X3)))), file('/home/ubuntu/library/tptp/Problems/COM/COM020+4.p', m__715)).
fof(f240, plain, sdtmndtasgtdt0(xu, xR, xb), inference(cnf_transformation, [], [f128])).
fof(f4495, plain, (~ spl35_27 | ~ spl35_60 | ~ spl35_128), inference(avatar_contradiction_clause, [], [f4494])).
fof(f4494, plain, ($false | (~ spl35_27 | ~ spl35_60 | ~ spl35_128)), inference(subsumption_resolution, [], [f4493, f259])).
fof(f4493, plain, (~ aElement0(xd) | (~ spl35_27 | ~ spl35_60 | ~ spl35_128)), inference(subsumption_resolution, [], [f4488, f664])).
fof(f4488, plain, (sP6(xb, xd) | ~ aElement0(xd) | (~ spl35_27 | ~ spl35_60 | ~ spl35_128)), inference(resolution, [], [f3859, f3657])).
fof(f3657, plain, (~ sP7(xd, xu) | ~ spl35_27), inference(backward_demodulation, [], [f625, f411])).
fof(f411, plain, ((xu = xw) | ~ spl35_27), inference(avatar_component_clause, [], [f409])).
fof(f409, plain, (spl35_27 <=> (xu = xw)), introduced(avatar_definition, [new_symbols(naming, [spl35_27])])).
fof(f625, plain, ~ sP7(xd, xw), inference(resolution, [], [f208, f264])).
fof(f3859, plain, (! [X0] : (sP7(X0, xu) | sP6(xb, X0) | ~ aElement0(X0)) | (~ spl35_27 | ~ spl35_60 | ~ spl35_128)), inference(subsumption_resolution, [], [f3858, f3792])).
fof(f3858, plain, (! [X0] : (sP6(xb, X0) | sP7(X0, xu) | ~ aElement0(X0) | ~ aElement0(xu)) | (~ spl35_27 | ~ spl35_60)), inference(subsumption_resolution, [], [f3857, f202])).
fof(f3857, plain, (! [X0] : (sP6(xb, X0) | sP7(X0, xu) | ~ aElement0(xb) | ~ aElement0(X0) | ~ aElement0(xu)) | (~ spl35_27 | ~ spl35_60)), inference(subsumption_resolution, [], [f3849, f3803])).
fof(f3803, plain, (iLess0(xu, xa) | (~ spl35_27 | ~ spl35_60)), inference(forward_demodulation, [], [f1265, f411])).
fof(f1265, plain, (iLess0(xw, xa) | ~ spl35_60), inference(avatar_component_clause, [], [f1264])).
fof(f1264, plain, (spl35_60 <=> iLess0(xw, xa)), introduced(avatar_definition, [new_symbols(naming, [spl35_60])])).
fof(f3797, plain, spl35_128, inference(avatar_split_clause, [], [f234, f3791])).
fof(f234, plain, aElement0(xu), inference(cnf_transformation, [], [f128])).
fof(f3643, plain, spl35_124, inference(avatar_split_clause, [], [f248, f3629])).
fof(f3618, plain, (spl35_60 | ~ spl35_22), inference(avatar_split_clause, [], [f3594, f385, f1264])).
fof(f3594, plain, (iLess0(xw, xa) | ~ spl35_22), inference(subsumption_resolution, [], [f2918, f201])).
fof(f2918, plain, (iLess0(xw, xa) | ~ aElement0(xa) | ~ spl35_22), inference(resolution, [], [f1023, f242])).
fof(f242, plain, aReductOfIn0(xv, xa, xR), inference(cnf_transformation, [], [f130])).
fof(f1023, plain, (! [X15] : (~ aReductOfIn0(xv, X15, xR) | iLess0(xw, X15) | ~ aElement0(X15)) | ~ spl35_22), inference(subsumption_resolution, [], [f1022, f248])).
fof(f1022, plain, (! [X15] : (iLess0(xw, X15) | ~ aReductOfIn0(xv, X15, xR) | ~ aElement0(xw) | ~ aElement0(X15)) | ~ spl35_22), inference(subsumption_resolution, [], [f1002, f241])).
fof(f1002, plain, (! [X15] : (iLess0(xw, X15) | ~ aReductOfIn0(xv, X15, xR) | ~ aElement0(xv) | ~ aElement0(xw) | ~ aElement0(X15)) | ~ spl35_22), inference(resolution, [], [f198, f387])).
fof(f387, plain, (sdtmndtplgtdt0(xv, xR, xw) | ~ spl35_22), inference(avatar_component_clause, [], [f385])).
fof(f198, plain, ! [X2, X0, X1] : (~ sdtmndtplgtdt0(X2, xR, X1) | iLess0(X1, X0) | ~ aReductOfIn0(X2, X0, xR) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f112])).
fof(f416, plain, (spl35_27 | spl35_28), inference(avatar_split_clause, [], [f252, f413, f409])).
fof(f252, plain, (sdtmndtplgtdt0(xu, xR, xw) | (xu = xw)), inference(cnf_transformation, [], [f133])).