fof(f19775, plain, $false, inference(avatar_sat_refutation, [], [f19767, f19770])).
fof(f19770, plain, ~ spl32_130, inference(avatar_split_clause, [], [f19769, f3076])).
fof(f3076, plain, (spl32_130 <=> sP8(sK23(xu, xv))), introduced(avatar_definition, [new_symbols(naming, [spl32_130])])).
fof(f19769, plain, ~ sP8(sK23(xu, xv)), inference(subsumption_resolution, [], [f19768, f224])).
fof(f224, plain, aElement0(xu), inference(cnf_transformation, [], [f123])).
fof(f123, plain, (sdtmndtasgtdt0(xu, xR, xb) & ((sdtmndtplgtdt0(xu, xR, xb) & ((sdtmndtplgtdt0(sK30, xR, xb) & aReductOfIn0(sK30, xu, xR) & aElement0(sK30)) | aReductOfIn0(xb, xu, xR))) | (xb = xu)) & aReductOfIn0(xu, xa, xR) & aElement0(xu)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK30])], [f20, f122])).
fof(f122, plain, (? [X0] : (sdtmndtplgtdt0(X0, xR, xb) & aReductOfIn0(X0, xu, xR) & aElement0(X0)) => (sdtmndtplgtdt0(sK30, xR, xb) & aReductOfIn0(sK30, xu, xR) & aElement0(sK30))), introduced(choice_axiom, [])).
fof(f20, plain, (sdtmndtasgtdt0(xu, xR, xb) & ((sdtmndtplgtdt0(xu, xR, xb) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xb) & aReductOfIn0(X0, xu, xR) & aElement0(X0)) | aReductOfIn0(xb, xu, xR))) | (xb = xu)) & aReductOfIn0(xu, xa, xR) & aElement0(xu)), file('/home/ubuntu/library/tptp/Problems/COM/COM017+4.p', m__755)).
fof(f19768, plain, (~ aElement0(xu) | ~ sP8(sK23(xu, xv))), inference(subsumption_resolution, [], [f3889, f225])).
fof(f225, plain, aReductOfIn0(xu, xa, xR), inference(cnf_transformation, [], [f123])).
fof(f3889, plain, (~ aReductOfIn0(xu, xa, xR) | ~ aElement0(xu) | ~ sP8(sK23(xu, xv))), inference(resolution, [], [f1041, f242])).
fof(f242, plain, ! [X0] : (~ sdtmndtasgtdt0(xu, xR, X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ! [X0] : ((~ sdtmndtasgtdt0(xu, xR, X0) & ~ sdtmndtplgtdt0(xu, xR, X0) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, X0) | ~ aReductOfIn0(X1, xu, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(X0, xu, xR) & ~ (xu = X0)) | ~ sP8(X0)), inference(rectify, [], [f126])).
fof(f126, plain, ! [X0] : ((~ sdtmndtasgtdt0(xu, xR, X0) & ~ sdtmndtplgtdt0(xu, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, xu, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, xu, xR) & ~ (xu = X0)) | ~ sP8(X0)), inference(nnf_transformation, [], [f69])).
fof(f69, plain, ! [X0] : ((~ sdtmndtasgtdt0(xu, xR, X0) & ~ sdtmndtplgtdt0(xu, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, xu, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, xu, xR) & ~ (xu = X0)) | ~ sP8(X0)), inference(usedef, [], [e69])).
fof(e69, plain, ! [X0] : (sP8(X0) <=> (~ sdtmndtasgtdt0(xu, xR, X0) & ~ sdtmndtplgtdt0(xu, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, xu, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, xu, xR) & ~ (xu = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f1041, plain, ! [X15] : (sdtmndtasgtdt0(X15, xR, sK23(X15, xv)) | ~ aReductOfIn0(X15, xa, xR) | ~ aElement0(X15)), inference(subsumption_resolution, [], [f1040, f191])).
fof(f191, plain, aElement0(xa), inference(cnf_transformation, [], [f17])).
fof(f17, plain, (aElement0(xc) & aElement0(xb) & aElement0(xa)), file('/home/ubuntu/library/tptp/Problems/COM/COM017+4.p', m__731)).
fof(f1040, plain, ! [X15] : (sdtmndtasgtdt0(X15, xR, sK23(X15, xv)) | ~ aReductOfIn0(X15, xa, xR) | ~ aElement0(X15) | ~ aElement0(xa)), inference(subsumption_resolution, [], [f1029, f231])).
fof(f231, plain, aElement0(xv), inference(cnf_transformation, [], [f125])).
fof(f125, plain, (sdtmndtasgtdt0(xv, xR, xc) & ((sdtmndtplgtdt0(xv, xR, xc) & ((sdtmndtplgtdt0(sK31, xR, xc) & aReductOfIn0(sK31, xv, xR) & aElement0(sK31)) | aReductOfIn0(xc, xv, xR))) | (xc = xv)) & aReductOfIn0(xv, xa, xR) & aElement0(xv)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK31])], [f21, f124])).
fof(f124, plain, (? [X0] : (sdtmndtplgtdt0(X0, xR, xc) & aReductOfIn0(X0, xv, xR) & aElement0(X0)) => (sdtmndtplgtdt0(sK31, xR, xc) & aReductOfIn0(sK31, xv, xR) & aElement0(sK31))), introduced(choice_axiom, [])).
fof(f21, plain, (sdtmndtasgtdt0(xv, xR, xc) & ((sdtmndtplgtdt0(xv, xR, xc) & (? [X0] : (sdtmndtplgtdt0(X0, xR, xc) & aReductOfIn0(X0, xv, xR) & aElement0(X0)) | aReductOfIn0(xc, xv, xR))) | (xc = xv)) & aReductOfIn0(xv, xa, xR) & aElement0(xv)), file('/home/ubuntu/library/tptp/Problems/COM/COM017+4.p', m__779)).
fof(f1029, plain, ! [X15] : (sdtmndtasgtdt0(X15, xR, sK23(X15, xv)) | ~ aReductOfIn0(X15, xa, xR) | ~ aElement0(xv) | ~ aElement0(X15) | ~ aElement0(xa)), inference(resolution, [], [f180, f232])).
fof(f232, plain, aReductOfIn0(xv, xa, xR), inference(cnf_transformation, [], [f125])).
fof(f180, plain, ! [X4, X5, X3] : (~ aReductOfIn0(X5, X3, xR) | sdtmndtasgtdt0(X4, xR, sK23(X4, X5)) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3)), inference(cnf_transformation, [], [f107])).
fof(f107, plain, (isTerminating0(xR) & ! [X0, X1] : (iLess0(X1, X0) | (~ sdtmndtplgtdt0(X0, xR, X1) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X1) | ~ aReductOfIn0(X2, X0, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X1, X0, xR)) | ~ aElement0(X1) | ~ aElement0(X0)) & isLocallyConfluent0(xR) & ! [X3, X4, X5] : ((sdtmndtasgtdt0(X5, xR, sK23(X4, X5)) & ((sdtmndtplgtdt0(X5, xR, sK23(X4, X5)) & ((sdtmndtplgtdt0(sK24(X4, X5), xR, sK23(X4, X5)) & aReductOfIn0(sK24(X4, X5), X5, xR) & aElement0(sK24(X4, X5))) | aReductOfIn0(sK23(X4, X5), X5, xR))) | (sK23(X4, X5) = X5)) & sdtmndtasgtdt0(X4, xR, sK23(X4, X5)) & sP4(sK23(X4, X5), X4) & aElement0(sK23(X4, X5))) | ~ aReductOfIn0(X5, X3, xR) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK23, sK24])], [f64, f106, f105])).
fof(f105, plain, ! [X5, X4] : (? [X6] : (sdtmndtasgtdt0(X5, xR, X6) & ((sdtmndtplgtdt0(X5, xR, X6) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X6) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(X6, X5, xR))) | (X5 = X6)) & sdtmndtasgtdt0(X4, xR, X6) & sP4(X6, X4) & aElement0(X6)) => (sdtmndtasgtdt0(X5, xR, sK23(X4, X5)) & ((sdtmndtplgtdt0(X5, xR, sK23(X4, X5)) & (? [X7] : (sdtmndtplgtdt0(X7, xR, sK23(X4, X5)) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(sK23(X4, X5), X5, xR))) | (sK23(X4, X5) = X5)) & sdtmndtasgtdt0(X4, xR, sK23(X4, X5)) & sP4(sK23(X4, X5), X4) & aElement0(sK23(X4, X5)))), introduced(choice_axiom, [])).
fof(f106, plain, ! [X5, X4] : (? [X7] : (sdtmndtplgtdt0(X7, xR, sK23(X4, X5)) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) => (sdtmndtplgtdt0(sK24(X4, X5), xR, sK23(X4, X5)) & aReductOfIn0(sK24(X4, X5), X5, xR) & aElement0(sK24(X4, X5)))), introduced(choice_axiom, [])).
fof(f64, plain, (isTerminating0(xR) & ! [X0, X1] : (iLess0(X1, X0) | (~ sdtmndtplgtdt0(X0, xR, X1) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X1) | ~ aReductOfIn0(X2, X0, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X1, X0, xR)) | ~ aElement0(X1) | ~ aElement0(X0)) & isLocallyConfluent0(xR) & ! [X3, X4, X5] : (? [X6] : (sdtmndtasgtdt0(X5, xR, X6) & ((sdtmndtplgtdt0(X5, xR, X6) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X6) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(X6, X5, xR))) | (X5 = X6)) & sdtmndtasgtdt0(X4, xR, X6) & sP4(X6, X4) & aElement0(X6)) | ~ aReductOfIn0(X5, X3, xR) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3))), inference(definition_folding, [], [f53, e63])).
fof(f63, plain, ! [X6, X4] : ((sdtmndtplgtdt0(X4, xR, X6) & (? [X8] : (sdtmndtplgtdt0(X8, xR, X6) & aReductOfIn0(X8, X4, xR) & aElement0(X8)) | aReductOfIn0(X6, X4, xR))) | (X4 = X6) | ~ sP4(X6, X4)), inference(usedef, [], [e63])).
fof(e63, plain, ! [X6, X4] : (sP4(X6, X4) <=> ((sdtmndtplgtdt0(X4, xR, X6) & (? [X8] : (sdtmndtplgtdt0(X8, xR, X6) & aReductOfIn0(X8, X4, xR) & aElement0(X8)) | aReductOfIn0(X6, X4, xR))) | (X4 = X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f53, plain, (isTerminating0(xR) & ! [X0, X1] : (iLess0(X1, X0) | (~ sdtmndtplgtdt0(X0, xR, X1) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X1) | ~ aReductOfIn0(X2, X0, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X1, X0, xR)) | ~ aElement0(X1) | ~ aElement0(X0)) & isLocallyConfluent0(xR) & ! [X3, X4, X5] : (? [X6] : (sdtmndtasgtdt0(X5, xR, X6) & ((sdtmndtplgtdt0(X5, xR, X6) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X6) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(X6, X5, xR))) | (X5 = X6)) & sdtmndtasgtdt0(X4, xR, X6) & ((sdtmndtplgtdt0(X4, xR, X6) & (? [X8] : (sdtmndtplgtdt0(X8, xR, X6) & aReductOfIn0(X8, X4, xR) & aElement0(X8)) | aReductOfIn0(X6, X4, xR))) | (X4 = X6)) & aElement0(X6)) | ~ aReductOfIn0(X5, X3, xR) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3))), inference(flattening, [], [f52])).
fof(f52, plain, (isTerminating0(xR) & ! [X0, X1] : ((iLess0(X1, X0) | (~ sdtmndtplgtdt0(X0, xR, X1) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X1) | ~ aReductOfIn0(X2, X0, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X1, X0, xR))) | (~ aElement0(X1) | ~ aElement0(X0))) & isLocallyConfluent0(xR) & ! [X3, X4, X5] : (? [X6] : (sdtmndtasgtdt0(X5, xR, X6) & ((sdtmndtplgtdt0(X5, xR, X6) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X6) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(X6, X5, xR))) | (X5 = X6)) & sdtmndtasgtdt0(X4, xR, X6) & ((sdtmndtplgtdt0(X4, xR, X6) & (? [X8] : (sdtmndtplgtdt0(X8, xR, X6) & aReductOfIn0(X8, X4, xR) & aElement0(X8)) | aReductOfIn0(X6, X4, xR))) | (X4 = X6)) & aElement0(X6)) | (~ aReductOfIn0(X5, X3, xR) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3)))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, (isTerminating0(xR) & ! [X0, X1] : ((aElement0(X1) & aElement0(X0)) => ((sdtmndtplgtdt0(X0, xR, X1) | ? [X2] : (sdtmndtplgtdt0(X2, xR, X1) & aReductOfIn0(X2, X0, xR) & aElement0(X2)) | aReductOfIn0(X1, X0, xR)) => iLess0(X1, X0))) & isLocallyConfluent0(xR) & ! [X3, X4, X5] : ((aReductOfIn0(X5, X3, xR) & aReductOfIn0(X4, X3, xR) & aElement0(X5) & aElement0(X4) & aElement0(X3)) => ? [X6] : (sdtmndtasgtdt0(X5, xR, X6) & ((sdtmndtplgtdt0(X5, xR, X6) & (? [X7] : (sdtmndtplgtdt0(X7, xR, X6) & aReductOfIn0(X7, X5, xR) & aElement0(X7)) | aReductOfIn0(X6, X5, xR))) | (X5 = X6)) & sdtmndtasgtdt0(X4, xR, X6) & ((sdtmndtplgtdt0(X4, xR, X6) & (? [X8] : (sdtmndtplgtdt0(X8, xR, X6) & aReductOfIn0(X8, X4, xR) & aElement0(X8)) | aReductOfIn0(X6, X4, xR))) | (X4 = X6)) & aElement0(X6)))), inference(rectify, [], [f16])).
fof(f16, plain, (isTerminating0(xR) & ! [X0, X1] : ((aElement0(X1) & aElement0(X0)) => ((sdtmndtplgtdt0(X0, xR, X1) | ? [X2] : (sdtmndtplgtdt0(X2, xR, X1) & aReductOfIn0(X2, X0, xR) & aElement0(X2)) | aReductOfIn0(X1, X0, xR)) => iLess0(X1, X0))) & isLocallyConfluent0(xR) & ! [X0, X1, X2] : ((aReductOfIn0(X2, X0, xR) & aReductOfIn0(X1, X0, xR) & aElement0(X2) & aElement0(X1) & aElement0(X0)) => ? [X3] : (sdtmndtasgtdt0(X2, xR, X3) & ((sdtmndtplgtdt0(X2, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X2, xR) & aElement0(X4)) | aReductOfIn0(X3, X2, xR))) | (X2 = X3)) & sdtmndtasgtdt0(X1, xR, X3) & ((sdtmndtplgtdt0(X1, xR, X3) & (? [X4] : (sdtmndtplgtdt0(X4, xR, X3) & aReductOfIn0(X4, X1, xR) & aElement0(X4)) | aReductOfIn0(X3, X1, xR))) | (X1 = X3)) & aElement0(X3)))), file('/home/ubuntu/library/tptp/Problems/COM/COM017+4.p', m__656_01)).
fof(f19767, plain, spl32_130, inference(avatar_split_clause, [], [f19744, f3076])).
fof(f19744, plain, sP8(sK23(xu, xv)), inference(subsumption_resolution, [], [f19724, f224])).
fof(f19724, plain, (~ aElement0(xu) | sP8(sK23(xu, xv))), inference(resolution, [], [f3930, f225])).
fof(f3930, plain, ! [X0] : (~ aReductOfIn0(X0, xa, xR) | ~ aElement0(X0) | sP8(sK23(X0, xv))), inference(subsumption_resolution, [], [f3921, f812])).
fof(f812, plain, ! [X0] : (~ aReductOfIn0(X0, xa, xR) | aElement0(sK23(X0, xv)) | ~ aElement0(X0)), inference(subsumption_resolution, [], [f811, f191])).
fof(f811, plain, ! [X0] : (aElement0(sK23(X0, xv)) | ~ aReductOfIn0(X0, xa, xR) | ~ aElement0(X0) | ~ aElement0(xa)), inference(subsumption_resolution, [], [f801, f231])).
fof(f801, plain, ! [X0] : (aElement0(sK23(X0, xv)) | ~ aReductOfIn0(X0, xa, xR) | ~ aElement0(xv) | ~ aElement0(X0) | ~ aElement0(xa)), inference(resolution, [], [f178, f232])).
fof(f178, plain, ! [X4, X5, X3] : (~ aReductOfIn0(X5, X3, xR) | aElement0(sK23(X4, X5)) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3)), inference(cnf_transformation, [], [f107])).
fof(f3921, plain, ! [X0] : (~ aReductOfIn0(X0, xa, xR) | ~ aElement0(X0) | sP8(sK23(X0, xv)) | ~ aElement0(sK23(X0, xv))), inference(resolution, [], [f1065, f247])).
fof(f247, plain, ! [X0] : (~ sdtmndtasgtdt0(xv, xR, X0) | sP8(X0) | ~ aElement0(X0)), inference(cnf_transformation, [], [f70])).
fof(f70, plain, ! [X0] : ((~ sdtmndtasgtdt0(xv, xR, X0) & ~ sdtmndtplgtdt0(xv, xR, X0) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, X0) | ~ aReductOfIn0(X1, xv, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(X0, xv, xR) & ~ (xv = X0)) | sP8(X0) | ~ aElement0(X0)), inference(definition_folding, [], [f56, e69])).
fof(f56, plain, ! [X0] : ((~ sdtmndtasgtdt0(xv, xR, X0) & ~ sdtmndtplgtdt0(xv, xR, X0) & ! [X1] : (~ sdtmndtplgtdt0(X1, xR, X0) | ~ aReductOfIn0(X1, xv, xR) | ~ aElement0(X1)) & ~ aReductOfIn0(X0, xv, xR) & ~ (xv = X0)) | (~ sdtmndtasgtdt0(xu, xR, X0) & ~ sdtmndtplgtdt0(xu, xR, X0) & ! [X2] : (~ sdtmndtplgtdt0(X2, xR, X0) | ~ aReductOfIn0(X2, xu, xR) | ~ aElement0(X2)) & ~ aReductOfIn0(X0, xu, xR) & ~ (xu = X0)) | ~ aElement0(X0)), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ~ ? [X0] : ((sdtmndtasgtdt0(xv, xR, X0) | sdtmndtplgtdt0(xv, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xv, xR) & aElement0(X1)) | aReductOfIn0(X0, xv, xR) | (xv = X0)) & (sdtmndtasgtdt0(xu, xR, X0) | sdtmndtplgtdt0(xu, xR, X0) | ? [X2] : (sdtmndtplgtdt0(X2, xR, X0) & aReductOfIn0(X2, xu, xR) & aElement0(X2)) | aReductOfIn0(X0, xu, xR) | (xu = X0)) & aElement0(X0)), inference(rectify, [], [f23])).
fof(f23, plain, ~ ? [X0] : ((sdtmndtasgtdt0(xv, xR, X0) | sdtmndtplgtdt0(xv, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xv, xR) & aElement0(X1)) | aReductOfIn0(X0, xv, xR) | (xv = X0)) & (sdtmndtasgtdt0(xu, xR, X0) | sdtmndtplgtdt0(xu, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xu, xR) & aElement0(X1)) | aReductOfIn0(X0, xu, xR) | (xu = X0)) & aElement0(X0)), inference(negated_conjecture, [], [f22])).
fof(f22, plain, ~ ? [X0] : ((sdtmndtasgtdt0(xv, xR, X0) | sdtmndtplgtdt0(xv, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xv, xR) & aElement0(X1)) | aReductOfIn0(X0, xv, xR) | (xv = X0)) & (sdtmndtasgtdt0(xu, xR, X0) | sdtmndtplgtdt0(xu, xR, X0) | ? [X1] : (sdtmndtplgtdt0(X1, xR, X0) & aReductOfIn0(X1, xu, xR) & aElement0(X1)) | aReductOfIn0(X0, xu, xR) | (xu = X0)) & aElement0(X0)), file('/home/ubuntu/library/tptp/Problems/COM/COM017+4.p', m__)).
fof(f1065, plain, ! [X15] : (sdtmndtasgtdt0(xv, xR, sK23(X15, xv)) | ~ aReductOfIn0(X15, xa, xR) | ~ aElement0(X15)), inference(subsumption_resolution, [], [f1064, f191])).
fof(f1064, plain, ! [X15] : (sdtmndtasgtdt0(xv, xR, sK23(X15, xv)) | ~ aReductOfIn0(X15, xa, xR) | ~ aElement0(X15) | ~ aElement0(xa)), inference(subsumption_resolution, [], [f1053, f231])).
fof(f1053, plain, ! [X15] : (sdtmndtasgtdt0(xv, xR, sK23(X15, xv)) | ~ aReductOfIn0(X15, xa, xR) | ~ aElement0(xv) | ~ aElement0(X15) | ~ aElement0(xa)), inference(resolution, [], [f185, f232])).
fof(f185, plain, ! [X4, X5, X3] : (~ aReductOfIn0(X5, X3, xR) | sdtmndtasgtdt0(X5, xR, sK23(X4, X5)) | ~ aReductOfIn0(X4, X3, xR) | ~ aElement0(X5) | ~ aElement0(X4) | ~ aElement0(X3)), inference(cnf_transformation, [], [f107])).