fof(f83194, plain, $false, inference(trivial_inequality_removal, [], [f82920])).
fof(f82920, plain, ~ (mult(a, mult(b, mult(a, c))) = mult(a, mult(b, mult(a, c)))), inference(superposition, [], [f45, f63811])).
fof(f63811, plain, ! [X4, X2, X3] : (mult(mult(X3, mult(X2, X3)), X4) = mult(X3, mult(X2, mult(X3, X4)))), inference(forward_demodulation, [], [f63810, f27])).
fof(f27, plain, ! [X0, X1] : (ld(X1, mult(X1, X0)) = X0), inference(cnf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1] : (ld(X1, mult(X1, X0)) = X0), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', f02)).
fof(f63810, plain, ! [X4, X2, X3] : (mult(X3, mult(X2, mult(X3, X4))) = mult(ld(X2, mult(X2, mult(X3, mult(X2, X3)))), X4)), inference(forward_demodulation, [], [f63644, f12540])).
fof(f12540, plain, ! [X12, X11] : (mult(X11, mult(X12, mult(X11, X12))) = mult(mult(X11, X12), mult(X11, X12))), inference(forward_demodulation, [], [f12477, f100])).
fof(f100, plain, ! [X6, X7, X5] : (mult(X5, mult(X6, mult(X7, X6))) = rd(mult(mult(X5, X6), mult(X7, mult(X6, X7))), X7)), inference(superposition, [], [f29, f40])).
fof(f40, plain, ! [X2, X0, X1] : (mult(mult(X2, X1), mult(X0, mult(X1, X0))) = mult(mult(X2, mult(X1, mult(X0, X1))), X0)), inference(forward_demodulation, [], [f39, f33])).
fof(f33, plain, ! [X0, X1] : (mult(mult(X1, X0), X1) = mult(X1, mult(X0, X1))), inference(cnf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1] : (mult(mult(X1, X0), X1) = mult(X1, mult(X0, X1))), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', f08)).
fof(f39, plain, ! [X2, X0, X1] : (mult(mult(X2, mult(mult(X1, X0), X1)), X0) = mult(mult(X2, X1), mult(X0, mult(X1, X0)))), inference(backward_demodulation, [], [f32, f33])).
fof(f32, plain, ! [X2, X0, X1] : (mult(mult(X2, X1), mult(mult(X0, X1), X0)) = mult(mult(X2, mult(mult(X1, X0), X1)), X0)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ! [X0, X1, X2] : (mult(mult(X2, X1), mult(mult(X0, X1), X0)) = mult(mult(X2, mult(mult(X1, X0), X1)), X0)), inference(rectify, [], [f7])).
fof(f7, plain, ! [X2, X0, X1] : (mult(mult(X1, X0), mult(mult(X2, X0), X2)) = mult(mult(X1, mult(mult(X0, X2), X0)), X2)), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', f07)).
fof(f29, plain, ! [X0, X1] : (rd(mult(X1, X0), X0) = X1), inference(cnf_transformation, [], [f4])).
fof(f4, plain, ! [X0, X1] : (rd(mult(X1, X0), X0) = X1), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', f04)).
fof(f12477, plain, ! [X12, X11] : (mult(mult(X11, X12), mult(X11, X12)) = rd(mult(mult(X11, X12), mult(X11, mult(X12, X11))), X11)), inference(superposition, [], [f366, f33])).
fof(f366, plain, ! [X12, X13] : (mult(X12, X12) = rd(mult(X12, mult(X12, X13)), X13)), inference(superposition, [], [f29, f146])).
fof(f146, plain, ! [X15, X16] : (mult(mult(X15, X15), X16) = mult(X15, mult(X15, X16))), inference(forward_demodulation, [], [f145, f31])).
fof(f31, plain, ! [X0] : (mult(unit, X0) = X0), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : (mult(unit, X0) = X0), inference(rectify, [], [f6])).
fof(f6, plain, ! [X1] : (mult(unit, X1) = X1), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', f06)).
fof(f145, plain, ! [X15, X16] : (mult(X15, mult(unit, mult(X15, X16))) = mult(mult(X15, X15), X16)), inference(subsumption_resolution, [], [f144, f31])).
fof(f144, plain, ! [X15, X16] : (~ (mult(X15, mult(X16, X15)) = mult(unit, mult(X15, mult(X16, X15)))) | (mult(X15, mult(unit, mult(X15, X16))) = mult(mult(X15, X15), X16))), inference(forward_demodulation, [], [f130, f33])).
fof(f130, plain, ! [X15, X16] : (~ (mult(unit, mult(X15, mult(X16, X15))) = mult(mult(X15, X16), X15)) | (mult(X15, mult(unit, mult(X15, X16))) = mult(mult(X15, X15), X16))), inference(superposition, [], [f41, f31])).
fof(f41, plain, ! [X2, X0, X1] : (~ (mult(X0, mult(X1, mult(X2, X1))) = mult(mult(mult(X0, X1), X2), X1)) | (mult(X1, mult(X0, mult(X1, X2))) = mult(mult(X1, mult(X0, X1)), X2))), inference(forward_demodulation, [], [f35, f33])).
fof(f35, plain, ! [X2, X0, X1] : ((mult(X1, mult(X0, mult(X1, X2))) = mult(mult(mult(X1, X0), X1), X2)) | ~ (mult(X0, mult(X1, mult(X2, X1))) = mult(mult(mult(X0, X1), X2), X1))), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ! [X0, X1, X2] : ((mult(X1, mult(X0, mult(X1, X2))) = mult(mult(mult(X1, X0), X1), X2)) | ~ (mult(X0, mult(X1, mult(X2, X1))) = mult(mult(mult(X0, X1), X2), X1))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ! [X0, X1, X2] : ((mult(X0, mult(X1, mult(X2, X1))) = mult(mult(mult(X0, X1), X2), X1)) => (mult(X1, mult(X0, mult(X1, X2))) = mult(mult(mult(X1, X0), X1), X2))), inference(rectify, [], [f10])).
fof(f10, plain, ! [X0, X1, X2] : ((mult(X0, mult(X1, mult(X2, X1))) = mult(mult(mult(X0, X1), X2), X1)) => (mult(X1, mult(X0, mult(X1, X2))) = mult(mult(mult(X1, X0), X1), X2))), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', f10)).
fof(f63644, plain, ! [X4, X2, X3] : (mult(X3, mult(X2, mult(X3, X4))) = mult(ld(X2, mult(mult(X2, X3), mult(X2, X3))), X4)), inference(superposition, [], [f62684, f27])).
fof(f62684, plain, ! [X4, X2, X3] : (mult(ld(X2, X3), mult(X2, mult(ld(X2, X3), X4))) = mult(ld(X2, mult(X3, X3)), X4)), inference(trivial_inequality_removal, [], [f62678])).
fof(f62678, plain, ! [X4, X2, X3] : (~ (mult(mult(X3, X4), ld(X2, X3)) = mult(mult(X3, X4), ld(X2, X3))) | (mult(ld(X2, X3), mult(X2, mult(ld(X2, X3), X4))) = mult(ld(X2, mult(X3, X3)), X4))), inference(backward_demodulation, [], [f27279, f62602])).
fof(f62602, plain, ! [X4, X2, X3] : (mult(X2, mult(ld(X2, X3), mult(X4, ld(X2, X3)))) = mult(mult(X3, X4), ld(X2, X3))), inference(backward_demodulation, [], [f1168, f62472])).
fof(f62472, plain, ! [X37, X38, X36] : (rd(mult(X36, mult(X38, mult(X37, X38))), X38) = mult(mult(X36, X38), X37)), inference(backward_demodulation, [], [f1179, f62471])).
fof(f62471, plain, ! [X4, X2, X3] : (mult(mult(X4, X3), X2) = mult(rd(X4, X2), mult(X2, mult(X3, X2)))), inference(forward_demodulation, [], [f62215, f33])).
fof(f62215, plain, ! [X4, X2, X3] : (mult(mult(X4, X3), X2) = mult(rd(X4, X2), mult(mult(X2, X3), X2))), inference(superposition, [], [f58170, f27])).
fof(f58170, plain, ! [X21, X19, X20] : (mult(mult(X20, ld(X21, X19)), X21) = mult(rd(X20, X21), mult(X19, X21))), inference(forward_demodulation, [], [f57950, f357])).
fof(f357, plain, ! [X4, X5] : (mult(X4, X5) = mult(f(X4), mult(f(X4), X5))), inference(superposition, [], [f146, f34])).
fof(f34, plain, ! [X0] : (mult(f(X0), f(X0)) = X0), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ! [X0] : (mult(f(X0), f(X0)) = X0), inference(rectify, [], [f9])).
fof(f9, plain, ! [X1] : (mult(f(X1), f(X1)) = X1), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', f09)).
fof(f57950, plain, ! [X21, X19, X20] : (mult(rd(X20, X21), mult(f(X19), mult(f(X19), X21))) = mult(mult(X20, ld(X21, X19)), X21)), inference(superposition, [], [f57712, f34])).
fof(f57712, plain, ! [X4, X2, X3] : (mult(mult(X4, ld(X2, mult(X3, X3))), X2) = mult(rd(X4, X2), mult(X3, mult(X3, X2)))), inference(backward_demodulation, [], [f27278, f57710])).
fof(f57710, plain, ! [X8, X7, X9] : (mult(mult(X9, ld(X8, X7)), mult(X7, X8)) = mult(rd(X9, X8), mult(X7, mult(X7, X8)))), inference(forward_demodulation, [], [f57687, f56201])).
fof(f56201, plain, ! [X163, X162] : (rd(X163, X162) = mult(X163, ld(X162, unit))), inference(forward_demodulation, [], [f56200, f54230])).
fof(f54230, plain, ! [X17, X16] : (mult(ld(X16, unit), mult(X16, X17)) = X17), inference(forward_demodulation, [], [f54229, f29])).
fof(f54229, plain, ! [X17, X16] : (mult(ld(X16, unit), mult(X16, X17)) = rd(mult(X17, rd(X17, X16)), rd(X17, X16))), inference(forward_demodulation, [], [f54055, f31])).
fof(f54055, plain, ! [X17, X16] : (mult(ld(X16, unit), mult(X16, X17)) = rd(mult(unit, mult(X17, rd(X17, X16))), rd(X17, X16))), inference(superposition, [], [f1263, f745])).
fof(f745, plain, ! [X0] : (unit = mult(ld(X0, unit), X0)), inference(superposition, [], [f639, f58])).
fof(f58, plain, ! [X2, X1] : (rd(X2, ld(X1, X2)) = X1), inference(superposition, [], [f29, f26])).
fof(f26, plain, ! [X0, X1] : (mult(X1, ld(X1, X0)) = X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0, X1] : (mult(X1, ld(X1, X0)) = X0), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', f01)).
fof(f639, plain, ! [X0] : (unit = mult(X0, rd(unit, X0))), inference(forward_demodulation, [], [f614, f59])).
fof(f59, plain, ! [X3] : (unit = rd(X3, X3)), inference(superposition, [], [f29, f31])).
fof(f614, plain, ! [X0] : (rd(X0, X0) = mult(X0, rd(unit, X0))), inference(superposition, [], [f271, f30])).
fof(f30, plain, ! [X0] : (mult(X0, unit) = X0), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ! [X0] : (mult(X0, unit) = X0), inference(rectify, [], [f5])).
fof(f5, plain, ! [X1] : (mult(X1, unit) = X1), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', f05)).
fof(f271, plain, ! [X8, X9] : (mult(X9, rd(X8, X9)) = rd(mult(X9, X8), X9)), inference(superposition, [], [f74, f28])).
fof(f28, plain, ! [X0, X1] : (mult(rd(X1, X0), X0) = X1), inference(cnf_transformation, [], [f3])).
fof(f3, plain, ! [X0, X1] : (mult(rd(X1, X0), X0) = X1), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', f03)).
fof(f74, plain, ! [X2, X1] : (mult(X1, X2) = rd(mult(X1, mult(X2, X1)), X1)), inference(superposition, [], [f29, f33])).
fof(f1263, plain, ! [X6, X8, X7] : (mult(X8, mult(X7, X6)) = rd(mult(mult(X8, X7), mult(X6, rd(X6, X7))), rd(X6, X7))), inference(forward_demodulation, [], [f1205, f28])).
fof(f1205, plain, ! [X6, X8, X7] : (mult(X8, mult(X7, mult(rd(X6, X7), X7))) = rd(mult(mult(X8, X7), mult(X6, rd(X6, X7))), rd(X6, X7))), inference(superposition, [], [f100, f69])).
fof(f69, plain, ! [X6, X5] : (mult(rd(X5, X6), mult(X6, rd(X5, X6))) = mult(X5, rd(X5, X6))), inference(superposition, [], [f33, f28])).
fof(f56200, plain, ! [X163, X162] : (mult(X163, ld(X162, unit)) = mult(ld(X162, unit), mult(X162, rd(X163, X162)))), inference(forward_demodulation, [], [f55060, f56109])).
fof(f56109, plain, ! [X10, X11] : (mult(mult(X10, X11), ld(X10, unit)) = mult(X10, rd(X11, X10))), inference(backward_demodulation, [], [f33966, f56104])).
fof(f56104, plain, ! [X239, X238] : (ld(X239, rd(X239, X238)) = ld(X238, unit)), inference(backward_demodulation, [], [f33963, f55059])).
fof(f55059, plain, ! [X161, X160] : (ld(X160, unit) = rd(X161, mult(X160, X161))), inference(superposition, [], [f29, f54230])).
fof(f33963, plain, ! [X239, X238] : (rd(X239, mult(X238, X239)) = ld(X239, rd(X239, X238))), inference(backward_demodulation, [], [f32986, f33962])).
fof(f33962, plain, ! [X48, X49] : (ld(mult(X48, X49), mult(X48, rd(X49, X48))) = ld(X49, rd(X49, X48))), inference(forward_demodulation, [], [f33961, f335])).
fof(f335, plain, ! [X8, X7] : (ld(mult(X7, X8), mult(X7, mult(X8, X8))) = X8), inference(superposition, [], [f27, f104])).
fof(f104, plain, ! [X0, X1] : (mult(mult(X1, X0), X0) = mult(X1, mult(X0, X0))), inference(forward_demodulation, [], [f103, f30])).
fof(f103, plain, ! [X0, X1] : (mult(mult(X1, unit), mult(X0, X0)) = mult(mult(X1, X0), X0)), inference(forward_demodulation, [], [f84, f31])).
fof(f84, plain, ! [X0, X1] : (mult(mult(X1, unit), mult(X0, mult(unit, X0))) = mult(mult(X1, mult(unit, X0)), X0)), inference(superposition, [], [f40, f30])).
fof(f33961, plain, ! [X48, X49] : (ld(mult(X48, X49), mult(X48, rd(X49, X48))) = ld(ld(mult(X48, X49), mult(X48, mult(X49, X49))), rd(X49, X48))), inference(forward_demodulation, [], [f33960, f24196])).
fof(f24196, plain, ! [X76, X77] : (rd(X76, X77) = ld(mult(X77, X76), mult(X77, mult(X76, rd(X76, X77))))), inference(forward_demodulation, [], [f24021, f69])).
fof(f24021, plain, ! [X76, X77] : (rd(X76, X77) = ld(mult(X77, X76), mult(X77, mult(rd(X76, X77), mult(X77, rd(X76, X77)))))), inference(superposition, [], [f680, f28])).
fof(f680, plain, ! [X8, X9] : (ld(mult(X8, mult(X9, X8)), mult(X8, mult(X9, mult(X8, X9)))) = X9), inference(superposition, [], [f27, f114])).
fof(f114, plain, ! [X10, X11] : (mult(mult(X10, mult(X11, X10)), X11) = mult(X10, mult(X11, mult(X10, X11)))), inference(forward_demodulation, [], [f95, f31])).
fof(f95, plain, ! [X10, X11] : (mult(mult(unit, X10), mult(X11, mult(X10, X11))) = mult(mult(X10, mult(X11, X10)), X11)), inference(superposition, [], [f40, f31])).
fof(f33960, plain, ! [X48, X49] : (ld(mult(X48, X49), mult(X48, rd(X49, X48))) = ld(ld(mult(X48, X49), mult(X48, mult(X49, X49))), ld(mult(X48, X49), mult(X48, mult(X49, rd(X49, X48)))))), inference(forward_demodulation, [], [f33959, f69])).
fof(f33959, plain, ! [X48, X49] : (ld(mult(X48, X49), mult(X48, rd(X49, X48))) = ld(ld(mult(X48, X49), mult(X48, mult(X49, X49))), ld(mult(X48, X49), mult(X48, mult(rd(X49, X48), mult(X48, rd(X49, X48))))))), inference(forward_demodulation, [], [f33805, f12540])).
fof(f33805, plain, ! [X48, X49] : (ld(mult(X48, X49), mult(X48, rd(X49, X48))) = ld(ld(mult(X48, X49), mult(X48, mult(X49, X49))), ld(mult(X48, X49), mult(mult(X48, rd(X49, X48)), mult(X48, rd(X49, X48)))))), inference(superposition, [], [f27282, f26227])).
fof(f26227, plain, ! [X35, X36] : (mult(mult(X36, rd(X35, X36)), mult(X36, X35)) = mult(X36, mult(X35, X35))), inference(backward_demodulation, [], [f2122, f26223])).
fof(f26223, plain, ! [X35, X34] : (mult(X35, X35) = mult(rd(X35, X34), mult(X34, X35))), inference(forward_demodulation, [], [f26222, f366])).
fof(f26222, plain, ! [X35, X34] : (mult(rd(X35, X34), mult(X34, X35)) = rd(mult(X35, mult(X35, rd(X35, X34))), rd(X35, X34))), inference(forward_demodulation, [], [f26039, f1796])).
fof(f1796, plain, ! [X4, X3] : (mult(X3, mult(X3, rd(X3, X4))) = mult(rd(X3, X4), mult(X4, mult(X3, rd(X3, X4))))), inference(forward_demodulation, [], [f1795, f703])).
fof(f703, plain, ! [X19, X18] : (mult(mult(X19, X18), rd(X18, X19)) = mult(X19, mult(X18, rd(X18, X19)))), inference(forward_demodulation, [], [f661, f69])).
fof(f661, plain, ! [X19, X18] : (mult(X19, mult(rd(X18, X19), mult(X19, rd(X18, X19)))) = mult(mult(X19, X18), rd(X18, X19))), inference(superposition, [], [f114, f28])).
fof(f1795, plain, ! [X4, X3] : (mult(rd(X3, X4), mult(mult(X4, X3), rd(X3, X4))) = mult(X3, mult(X3, rd(X3, X4)))), inference(forward_demodulation, [], [f1727, f28])).
fof(f1727, plain, ! [X4, X3] : (mult(rd(X3, X4), mult(mult(X4, X3), rd(X3, X4))) = mult(mult(rd(X3, X4), X4), mult(X3, rd(X3, X4)))), inference(superposition, [], [f106, f33])).
fof(f106, plain, ! [X14, X12, X13] : (mult(mult(X14, mult(X13, X12)), rd(X12, X13)) = mult(mult(X14, X13), mult(X12, rd(X12, X13)))), inference(forward_demodulation, [], [f88, f69])).
fof(f88, plain, ! [X14, X12, X13] : (mult(mult(X14, X13), mult(rd(X12, X13), mult(X13, rd(X12, X13)))) = mult(mult(X14, mult(X13, X12)), rd(X12, X13))), inference(superposition, [], [f40, f28])).
fof(f26039, plain, ! [X35, X34] : (mult(rd(X35, X34), mult(X34, X35)) = rd(mult(rd(X35, X34), mult(X34, mult(X35, rd(X35, X34)))), rd(X35, X34))), inference(superposition, [], [f74, f703])).
fof(f2122, plain, ! [X35, X36] : (mult(mult(X36, rd(X35, X36)), mult(X36, X35)) = mult(X36, mult(rd(X35, X36), mult(X36, X35)))), inference(superposition, [], [f115, f28])).
fof(f115, plain, ! [X0, X1] : (mult(mult(X0, X1), mult(X0, mult(X1, X0))) = mult(X0, mult(X1, mult(X0, mult(X1, X0))))), inference(forward_demodulation, [], [f96, f114])).
fof(f96, plain, ! [X0, X1] : (mult(X0, mult(mult(X1, mult(X0, X1)), X0)) = mult(mult(X0, X1), mult(X0, mult(X1, X0)))), inference(superposition, [], [f40, f33])).
fof(f27282, plain, ! [X2, X1] : (ld(X1, X2) = ld(ld(X1, mult(X2, X1)), ld(X1, mult(X2, X2)))), inference(backward_demodulation, [], [f317, f27277])).
fof(f27277, plain, ! [X6, X7] : (mult(ld(X6, X7), X7) = ld(X6, mult(X7, X7))), inference(forward_demodulation, [], [f27182, f14623])).
fof(f14623, plain, ! [X14, X13] : (ld(X14, mult(X13, X13)) = rd(ld(X14, mult(X13, mult(X13, X14))), X14)), inference(superposition, [], [f541, f146])).
fof(f541, plain, ! [X2, X3] : (ld(X2, X3) = rd(ld(X2, mult(X3, X2)), X2)), inference(superposition, [], [f29, f227])).
fof(f227, plain, ! [X4, X3] : (mult(ld(X3, X4), X3) = ld(X3, mult(X4, X3))), inference(superposition, [], [f27, f67])).
fof(f67, plain, ! [X2, X1] : (mult(X2, X1) = mult(X1, mult(ld(X1, X2), X1))), inference(superposition, [], [f33, f26])).
fof(f27182, plain, ! [X6, X7] : (mult(ld(X6, X7), X7) = rd(ld(X6, mult(X7, mult(X7, X6))), X6)), inference(backward_demodulation, [], [f24626, f26953])).
fof(f26953, plain, ! [X10, X11] : (mult(ld(X10, X11), mult(X11, X10)) = ld(X10, mult(X11, mult(X11, X10)))), inference(superposition, [], [f27, f1657])).
fof(f1657, plain, ! [X0, X1] : (mult(X1, mult(X1, X0)) = mult(X0, mult(ld(X0, X1), mult(X1, X0)))), inference(forward_demodulation, [], [f1656, f698])).
fof(f698, plain, ! [X2, X1] : (mult(mult(ld(X1, X2), X2), X1) = mult(ld(X1, X2), mult(X2, X1))), inference(forward_demodulation, [], [f697, f26])).
fof(f697, plain, ! [X2, X1] : (mult(mult(ld(X1, X2), X2), X1) = mult(ld(X1, X2), mult(X1, ld(X1, mult(X2, X1))))), inference(forward_demodulation, [], [f653, f227])).
fof(f653, plain, ! [X2, X1] : (mult(ld(X1, X2), mult(X1, mult(ld(X1, X2), X1))) = mult(mult(ld(X1, X2), X2), X1)), inference(superposition, [], [f114, f26])).
fof(f1656, plain, ! [X0, X1] : (mult(X1, mult(X1, X0)) = mult(X0, mult(mult(ld(X0, X1), X1), X0))), inference(forward_demodulation, [], [f1599, f26])).
fof(f1599, plain, ! [X0, X1] : (mult(X0, mult(mult(ld(X0, X1), X1), X0)) = mult(mult(X0, ld(X0, X1)), mult(X1, X0))), inference(superposition, [], [f105, f33])).
fof(f105, plain, ! [X4, X2, X3] : (mult(mult(X4, mult(ld(X2, X3), X3)), X2) = mult(mult(X4, ld(X2, X3)), mult(X3, X2))), inference(forward_demodulation, [], [f85, f67])).
fof(f85, plain, ! [X4, X2, X3] : (mult(mult(X4, ld(X2, X3)), mult(X2, mult(ld(X2, X3), X2))) = mult(mult(X4, mult(ld(X2, X3), X3)), X2)), inference(superposition, [], [f40, f26])).
fof(f24626, plain, ! [X6, X7] : (mult(ld(X6, X7), X7) = rd(mult(ld(X6, X7), mult(X7, X6)), X6)), inference(forward_demodulation, [], [f24625, f26])).
fof(f24625, plain, ! [X6, X7] : (mult(ld(X6, X7), mult(X6, ld(X6, X7))) = rd(mult(ld(X6, X7), mult(X7, X6)), X6)), inference(forward_demodulation, [], [f24533, f26])).
fof(f24533, plain, ! [X6, X7] : (mult(ld(X6, X7), mult(X6, ld(X6, X7))) = rd(mult(ld(X6, X7), mult(X6, ld(X6, mult(X7, X6)))), X6)), inference(superposition, [], [f681, f227])).
fof(f681, plain, ! [X10, X11] : (mult(X10, mult(X11, X10)) = rd(mult(X10, mult(X11, mult(X10, X11))), X11)), inference(superposition, [], [f29, f114])).
fof(f317, plain, ! [X2, X1] : (ld(X1, X2) = ld(ld(X1, mult(X2, X1)), mult(ld(X1, X2), X2))), inference(forward_demodulation, [], [f300, f227])).
fof(f300, plain, ! [X2, X1] : (ld(X1, X2) = ld(mult(ld(X1, X2), X1), mult(ld(X1, X2), X2))), inference(superposition, [], [f75, f26])).
fof(f75, plain, ! [X4, X3] : (ld(mult(X3, X4), mult(X3, mult(X4, X3))) = X3), inference(superposition, [], [f27, f33])).
fof(f32986, plain, ! [X239, X238] : (ld(mult(X238, X239), mult(X238, rd(X239, X238))) = rd(X239, mult(X238, X239))), inference(forward_demodulation, [], [f32425, f335])).
fof(f32425, plain, ! [X239, X238] : (ld(mult(X238, X239), mult(X238, rd(X239, X238))) = rd(ld(mult(X238, X239), mult(X238, mult(X239, X239))), mult(X238, X239))), inference(superposition, [], [f541, f26227])).
fof(f33966, plain, ! [X10, X11] : (mult(X10, rd(X11, X10)) = mult(mult(X10, X11), ld(X11, rd(X11, X10)))), inference(backward_demodulation, [], [f32874, f33963])).
fof(f32874, plain, ! [X10, X11] : (mult(mult(X10, X11), rd(X11, mult(X10, X11))) = mult(X10, rd(X11, X10))), inference(backward_demodulation, [], [f619, f32359])).
fof(f32359, plain, ! [X50, X51] : (mult(X50, rd(X51, X50)) = rd(mult(X50, mult(X51, X51)), mult(X50, X51))), inference(superposition, [], [f29, f26227])).
fof(f619, plain, ! [X10, X11] : (mult(mult(X10, X11), rd(X11, mult(X10, X11))) = rd(mult(X10, mult(X11, X11)), mult(X10, X11))), inference(superposition, [], [f271, f104])).
fof(f55060, plain, ! [X163, X162] : (mult(ld(X162, unit), mult(mult(X162, X163), ld(X162, unit))) = mult(X163, ld(X162, unit))), inference(superposition, [], [f33, f54230])).
fof(f57687, plain, ! [X8, X7, X9] : (mult(mult(X9, ld(X8, X7)), mult(X7, X8)) = mult(mult(X9, ld(X8, unit)), mult(X7, mult(X7, X8)))), inference(backward_demodulation, [], [f33096, f57476])).
fof(f57476, plain, ! [X39, X38] : (ld(X39, unit) = ld(mult(X38, X39), X38)), inference(superposition, [], [f56104, f29])).
fof(f33096, plain, ! [X8, X7, X9] : (mult(mult(X9, ld(mult(X7, X8), X7)), mult(X7, mult(X7, X8))) = mult(mult(X9, ld(X8, X7)), mult(X7, X8))), inference(superposition, [], [f27278, f25585])).
fof(f25585, plain, ! [X6, X7] : (ld(X6, X7) = ld(mult(X7, X6), mult(X7, X7))), inference(backward_demodulation, [], [f24138, f25574])).
fof(f25574, plain, ! [X66, X67] : (mult(X67, X67) = mult(X66, mult(ld(X66, X67), X67))), inference(forward_demodulation, [], [f25573, f366])).
fof(f25573, plain, ! [X66, X67] : (mult(X66, mult(ld(X66, X67), X67)) = rd(mult(X67, mult(X67, X66)), X66)), inference(forward_demodulation, [], [f25371, f1657])).
fof(f25371, plain, ! [X66, X67] : (mult(X66, mult(ld(X66, X67), X67)) = rd(mult(X66, mult(ld(X66, X67), mult(X67, X66))), X66)), inference(superposition, [], [f74, f698])).
fof(f24138, plain, ! [X6, X7] : (ld(X6, X7) = ld(mult(X7, X6), mult(X6, mult(ld(X6, X7), X7)))), inference(forward_demodulation, [], [f24137, f26])).
fof(f24137, plain, ! [X6, X7] : (ld(X6, X7) = ld(mult(X6, ld(X6, mult(X7, X6))), mult(X6, mult(ld(X6, X7), X7)))), inference(forward_demodulation, [], [f23987, f26])).
fof(f23987, plain, ! [X6, X7] : (ld(X6, X7) = ld(mult(X6, ld(X6, mult(X7, X6))), mult(X6, mult(ld(X6, X7), mult(X6, ld(X6, X7)))))), inference(superposition, [], [f680, f227])).
fof(f27278, plain, ! [X4, X2, X3] : (mult(mult(X4, ld(X2, X3)), mult(X3, X2)) = mult(mult(X4, ld(X2, mult(X3, X3))), X2)), inference(backward_demodulation, [], [f105, f27277])).
fof(f1179, plain, ! [X37, X38, X36] : (mult(rd(X36, X37), mult(X37, mult(X38, X37))) = rd(mult(X36, mult(X38, mult(X37, X38))), X38)), inference(superposition, [], [f100, f28])).
fof(f1168, plain, ! [X4, X2, X3] : (mult(X2, mult(ld(X2, X3), mult(X4, ld(X2, X3)))) = rd(mult(X3, mult(X4, mult(ld(X2, X3), X4))), X4)), inference(superposition, [], [f100, f26])).
fof(f27279, plain, ! [X4, X2, X3] : (~ (mult(X2, mult(ld(X2, X3), mult(X4, ld(X2, X3)))) = mult(mult(X3, X4), ld(X2, X3))) | (mult(ld(X2, X3), mult(X2, mult(ld(X2, X3), X4))) = mult(ld(X2, mult(X3, X3)), X4))), inference(backward_demodulation, [], [f126, f27277])).
fof(f126, plain, ! [X4, X2, X3] : (~ (mult(X2, mult(ld(X2, X3), mult(X4, ld(X2, X3)))) = mult(mult(X3, X4), ld(X2, X3))) | (mult(ld(X2, X3), mult(X2, mult(ld(X2, X3), X4))) = mult(mult(ld(X2, X3), X3), X4))), inference(superposition, [], [f41, f26])).
fof(f45, plain, ~ (mult(a, mult(b, mult(a, c))) = mult(mult(a, mult(b, a)), c)), inference(forward_demodulation, [], [f38, f33])).
fof(f38, plain, ~ (mult(a, mult(b, mult(a, c))) = mult(mult(mult(a, b), a), c)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ~ (mult(a, mult(b, mult(a, c))) = mult(mult(mult(a, b), a), c)), inference(flattening, [], [f14])).
fof(f14, plain, ~ (mult(a, mult(b, mult(a, c))) = mult(mult(mult(a, b), a), c)), inference(negated_conjecture, [], [f13])).
fof(f13, plain, ~ (mult(a, mult(b, mult(a, c))) = mult(mult(mult(a, b), a), c)), file('/home/ubuntu/library/tptp/Problems/GRP/GRP667+1.p', goals)).