fof(f175491, plain, $false, inference(avatar_sat_refutation, [], [f1713, f1739, f1744, f1760, f1767, f174441, f174598, f175490])).
fof(f175490, plain, (~ spl37_37 | ~ spl37_4450), inference(avatar_contradiction_clause, [], [f175489])).
fof(f175489, plain, ($false | (~ spl37_37 | ~ spl37_4450)), inference(subsumption_resolution, [], [f175488, f494])).
fof(f494, plain, ~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK36)), inference(cnf_transformation, [], [f287])).
fof(f287, plain, (((((~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK36)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK35)) & ! [X5] : (p(ap(sK35, X5)) | ~ p(ap(sK36, X5)) | ~ mem(X5, sK32)) & mem(sK36, arr(sK32, bool))) & mem(sK35, arr(sK32, bool))) & mem(sK34, arr(sK33, sK32))) & ne(sK33)) & ne(sK32)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK32, sK33, sK34, sK35, sK36])], [f156, f286, f285, f284, f283, f282])).
fof(f282, plain, (? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, X0), X2), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, X0), X2), X3)) & ! [X5] : (p(ap(X3, X5)) | ~ p(ap(X4, X5)) | ~ mem(X5, X0)) & mem(X4, arr(X0, bool))) & mem(X3, arr(X0, bool))) & mem(X2, arr(X1, X0))) & ne(X1)) & ne(X0)) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, sK32), X2), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, sK32), X2), X3)) & ! [X5] : (p(ap(X3, X5)) | ~ p(ap(X4, X5)) | ~ mem(X5, sK32)) & mem(X4, arr(sK32, bool))) & mem(X3, arr(sK32, bool))) & mem(X2, arr(X1, sK32))) & ne(X1)) & ne(sK32))), introduced(choice_axiom, [])).
fof(f283, plain, (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, sK32), X2), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, sK32), X2), X3)) & ! [X5] : (p(ap(X3, X5)) | ~ p(ap(X4, X5)) | ~ mem(X5, sK32)) & mem(X4, arr(sK32, bool))) & mem(X3, arr(sK32, bool))) & mem(X2, arr(X1, sK32))) & ne(X1)) => (? [X2] : (? [X3] : (? [X4] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), X2), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), X2), X3)) & ! [X5] : (p(ap(X3, X5)) | ~ p(ap(X4, X5)) | ~ mem(X5, sK32)) & mem(X4, arr(sK32, bool))) & mem(X3, arr(sK32, bool))) & mem(X2, arr(sK33, sK32))) & ne(sK33))), introduced(choice_axiom, [])).
fof(f284, plain, (? [X2] : (? [X3] : (? [X4] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), X2), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), X2), X3)) & ! [X5] : (p(ap(X3, X5)) | ~ p(ap(X4, X5)) | ~ mem(X5, sK32)) & mem(X4, arr(sK32, bool))) & mem(X3, arr(sK32, bool))) & mem(X2, arr(sK33, sK32))) => (? [X3] : (? [X4] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), X3)) & ! [X5] : (p(ap(X3, X5)) | ~ p(ap(X4, X5)) | ~ mem(X5, sK32)) & mem(X4, arr(sK32, bool))) & mem(X3, arr(sK32, bool))) & mem(sK34, arr(sK33, sK32)))), introduced(choice_axiom, [])).
fof(f285, plain, (? [X3] : (? [X4] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), X3)) & ! [X5] : (p(ap(X3, X5)) | ~ p(ap(X4, X5)) | ~ mem(X5, sK32)) & mem(X4, arr(sK32, bool))) & mem(X3, arr(sK32, bool))) => (? [X4] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK35)) & ! [X5] : (p(ap(sK35, X5)) | ~ p(ap(X4, X5)) | ~ mem(X5, sK32)) & mem(X4, arr(sK32, bool))) & mem(sK35, arr(sK32, bool)))), introduced(choice_axiom, [])).
fof(f286, plain, (? [X4] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK35)) & ! [X5] : (p(ap(sK35, X5)) | ~ p(ap(X4, X5)) | ~ mem(X5, sK32)) & mem(X4, arr(sK32, bool))) => (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK36)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK35)) & ! [X5] : (p(ap(sK35, X5)) | ~ p(ap(sK36, X5)) | ~ mem(X5, sK32)) & mem(sK36, arr(sK32, bool)))), introduced(choice_axiom, [])).
fof(f156, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, X0), X2), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, X0), X2), X3)) & ! [X5] : (p(ap(X3, X5)) | ~ p(ap(X4, X5)) | ~ mem(X5, X0)) & mem(X4, arr(X0, bool))) & mem(X3, arr(X0, bool))) & mem(X2, arr(X1, X0))) & ne(X1)) & ne(X0)), inference(flattening, [], [f155])).
fof(f155, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (((~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, X0), X2), X4)) & p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, X0), X2), X3))) & ! [X5] : ((p(ap(X3, X5)) | ~ p(ap(X4, X5))) | ~ mem(X5, X0))) & mem(X4, arr(X0, bool))) & mem(X3, arr(X0, bool))) & mem(X2, arr(X1, X0))) & ne(X1)) & ne(X0)), inference(ennf_transformation, [], [f105])).
fof(f105, plain, ~ ! [X0] : (ne(X0) => ! [X1] : (ne(X1) => ! [X2] : (mem(X2, arr(X1, X0)) => ! [X3] : (mem(X3, arr(X0, bool)) => ! [X4] : (mem(X4, arr(X0, bool)) => (! [X5] : (mem(X5, X0) => (p(ap(X4, X5)) => p(ap(X3, X5)))) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, X0), X2), X3)) => p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X1, X0), X2), X4))))))))), inference(rectify, [], [f53])).
fof(f53, plain, ~ ! [X8] : (ne(X8) => ! [X9] : (ne(X9) => ! [X19] : (mem(X19, arr(X9, X8)) => ! [X20] : (mem(X20, arr(X8, bool)) => ! [X44] : (mem(X44, arr(X8, bool)) => (! [X45] : (mem(X45, X8) => (p(ap(X44, X45)) => p(ap(X20, X45)))) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X9, X8), X19), X20)) => p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X9, X8), X19), X44))))))))), inference(negated_conjecture, [], [f52])).
fof(f52, plain, ~ ! [X8] : (ne(X8) => ! [X9] : (ne(X9) => ! [X19] : (mem(X19, arr(X9, X8)) => ! [X20] : (mem(X20, arr(X8, bool)) => ! [X44] : (mem(X44, arr(X8, bool)) => (! [X45] : (mem(X45, X8) => (p(ap(X44, X45)) => p(ap(X20, X45)))) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X9, X8), X19), X20)) => p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X9, X8), X19), X44))))))))), file('/home/ubuntu/library/tptp/Problems/ITP/ITP006+2.p', conj_thm_2EquantHeuristics_2EGUESS__RULES__WEAKEN__FORALL__POINT)).
fof(f175488, plain, (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK36)) | (~ spl37_37 | ~ spl37_4450)), inference(subsumption_resolution, [], [f175487, f489])).
fof(f489, plain, mem(sK34, arr(sK33, sK32)), inference(cnf_transformation, [], [f287])).
fof(f175487, plain, (~ mem(sK34, arr(sK33, sK32)) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK36)) | (~ spl37_37 | ~ spl37_4450)), inference(subsumption_resolution, [], [f175474, f311])).
fof(f311, plain, ~ p(c_2Ebool_2EF), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ~ p(c_2Ebool_2EF), file('/home/ubuntu/library/tptp/Problems/ITP/ITP006+2.p', ax_false_p)).
fof(f175474, plain, (p(c_2Ebool_2EF) | ~ mem(sK34, arr(sK33, sK32)) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK36)) | (~ spl37_37 | ~ spl37_4450)), inference(superposition, [], [f5284, f174436])).
fof(f174436, plain, ((c_2Ebool_2EF = ap(sK36, ap(sK34, sK31(sK33, sK34, sK36)))) | ~ spl37_4450), inference(avatar_component_clause, [], [f174434])).
fof(f174434, plain, (spl37_4450 <=> (c_2Ebool_2EF = ap(sK36, ap(sK34, sK31(sK33, sK34, sK36))))), introduced(avatar_definition, [new_symbols(naming, [spl37_4450])])).
fof(f5284, plain, (! [X0] : (p(ap(sK36, ap(X0, sK31(sK33, X0, sK36)))) | ~ mem(X0, arr(sK33, sK32)) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), X0), sK36))) | ~ spl37_37), inference(subsumption_resolution, [], [f5278, f488])).
fof(f488, plain, ne(sK33), inference(cnf_transformation, [], [f287])).
fof(f5278, plain, (! [X0] : (~ ne(sK33) | p(ap(sK36, ap(X0, sK31(sK33, X0, sK36)))) | ~ mem(X0, arr(sK33, sK32)) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), X0), sK36))) | ~ spl37_37), inference(resolution, [], [f1743, f489])).
fof(f1743, plain, (! [X33, X36, X34] : (~ mem(X36, arr(X34, sK32)) | ~ ne(X34) | p(ap(sK36, ap(X33, sK31(X34, X33, sK36)))) | ~ mem(X33, arr(X34, sK32)) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X34, sK32), X33), sK36))) | ~ spl37_37), inference(avatar_component_clause, [], [f1742])).
fof(f1742, plain, (spl37_37 <=> ! [X34, X36, X33] : (p(ap(sK36, ap(X33, sK31(X34, X33, sK36)))) | ~ ne(X34) | ~ mem(X36, arr(X34, sK32)) | ~ mem(X33, arr(X34, sK32)) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X34, sK32), X33), sK36)))), introduced(avatar_definition, [new_symbols(naming, [spl37_37])])).
fof(f174598, plain, (~ spl37_36 | spl37_4451), inference(avatar_contradiction_clause, [], [f174597])).
fof(f174597, plain, ($false | (~ spl37_36 | spl37_4451)), inference(subsumption_resolution, [], [f174596, f21993])).
fof(f21993, plain, (mem(sK31(sK33, sK34, sK36), sK33) | ~ spl37_36), inference(subsumption_resolution, [], [f21990, f489])).
fof(f21990, plain, (~ mem(sK34, arr(sK33, sK32)) | mem(sK31(sK33, sK34, sK36), sK33) | ~ spl37_36), inference(resolution, [], [f20586, f494])).
fof(f20586, plain, (! [X18] : (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), X18), sK36)) | ~ mem(X18, arr(sK33, sK32)) | mem(sK31(sK33, X18, sK36), sK33)) | ~ spl37_36), inference(subsumption_resolution, [], [f20278, f488])).
fof(f20278, plain, (! [X18] : (~ ne(sK33) | mem(sK31(sK33, X18, sK36), sK33) | ~ mem(X18, arr(sK33, sK32)) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), X18), sK36))) | ~ spl37_36), inference(resolution, [], [f1712, f489])).
fof(f1712, plain, (! [X33, X36, X34] : (~ mem(X36, arr(X33, sK32)) | ~ ne(X33) | mem(sK31(X33, X34, sK36), X33) | ~ mem(X34, arr(X33, sK32)) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X33, sK32), X34), sK36))) | ~ spl37_36), inference(avatar_component_clause, [], [f1711])).
fof(f1711, plain, (spl37_36 <=> ! [X34, X36, X33] : (mem(sK31(X33, X34, sK36), X33) | ~ ne(X33) | ~ mem(X36, arr(X33, sK32)) | ~ mem(X34, arr(X33, sK32)) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X33, sK32), X34), sK36)))), introduced(avatar_definition, [new_symbols(naming, [spl37_36])])).
fof(f174596, plain, (~ mem(sK31(sK33, sK34, sK36), sK33) | spl37_4451), inference(resolution, [], [f174440, f670])).
fof(f670, plain, ! [X10] : (mem(ap(sK34, X10), sK32) | ~ mem(X10, sK33)), inference(resolution, [], [f291, f489])).
fof(f291, plain, ! [X2, X0, X3, X1] : (~ mem(X2, arr(X0, X1)) | ~ mem(X3, X0) | mem(ap(X2, X3), X1)), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ! [X0, X1, X2] : (! [X3] : (mem(ap(X2, X3), X1) | ~ mem(X3, X0)) | ~ mem(X2, arr(X0, X1))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ! [X0, X1, X2] : (mem(X2, arr(X0, X1)) => ! [X3] : (mem(X3, X0) => mem(ap(X2, X3), X1))), file('/home/ubuntu/library/tptp/Problems/ITP/ITP006+2.p', ap_tp)).
fof(f174440, plain, (~ mem(ap(sK34, sK31(sK33, sK34, sK36)), sK32) | spl37_4451), inference(avatar_component_clause, [], [f174438])).
fof(f174438, plain, (spl37_4451 <=> mem(ap(sK34, sK31(sK33, sK34, sK36)), sK32)), introduced(avatar_definition, [new_symbols(naming, [spl37_4451])])).
fof(f174441, plain, (spl37_4450 | ~ spl37_4451 | ~ spl37_36 | ~ spl37_40), inference(avatar_split_clause, [], [f174432, f1758, f1711, f174438, f174434])).
fof(f1758, plain, (spl37_40 <=> ! [X0] : (~ mem(X0, sK33) | ~ p(ap(sK35, ap(sK34, X0))))), introduced(avatar_definition, [new_symbols(naming, [spl37_40])])).
fof(f174432, plain, (~ mem(ap(sK34, sK31(sK33, sK34, sK36)), sK32) | (c_2Ebool_2EF = ap(sK36, ap(sK34, sK31(sK33, sK34, sK36)))) | (~ spl37_36 | ~ spl37_40)), inference(subsumption_resolution, [], [f174376, f311])).
fof(f174376, plain, (p(c_2Ebool_2EF) | ~ mem(ap(sK34, sK31(sK33, sK34, sK36)), sK32) | (c_2Ebool_2EF = ap(sK36, ap(sK34, sK31(sK33, sK34, sK36)))) | (~ spl37_36 | ~ spl37_40)), inference(superposition, [], [f797, f24845])).
fof(f24845, plain, ((c_2Ebool_2EF = ap(sK35, ap(sK34, sK31(sK33, sK34, sK36)))) | (~ spl37_36 | ~ spl37_40)), inference(resolution, [], [f4924, f21993])).
fof(f4924, plain, (! [X1] : (~ mem(X1, sK33) | (c_2Ebool_2EF = ap(sK35, ap(sK34, X1)))) | ~ spl37_40), inference(subsumption_resolution, [], [f1802, f670])).
fof(f1802, plain, (! [X1] : (~ mem(X1, sK33) | (c_2Ebool_2EF = ap(sK35, ap(sK34, X1))) | ~ mem(ap(sK34, X1), sK32)) | ~ spl37_40), inference(resolution, [], [f1759, f711])).
fof(f711, plain, ! [X0] : (p(ap(sK35, X0)) | (c_2Ebool_2EF = ap(sK35, X0)) | ~ mem(X0, sK32)), inference(resolution, [], [f677, f671])).
fof(f671, plain, ! [X11] : (mem(ap(sK35, X11), bool) | ~ mem(X11, sK32)), inference(resolution, [], [f291, f490])).
fof(f490, plain, mem(sK35, arr(sK32, bool)), inference(cnf_transformation, [], [f287])).
fof(f677, plain, ! [X1] : (~ mem(X1, bool) | (c_2Ebool_2EF = X1) | p(X1)), inference(subsumption_resolution, [], [f676, f311])).
fof(f676, plain, ! [X1] : (p(c_2Ebool_2EF) | p(X1) | (c_2Ebool_2EF = X1) | ~ mem(X1, bool)), inference(resolution, [], [f292, f310])).
fof(f310, plain, mem(c_2Ebool_2EF, bool), inference(cnf_transformation, [], [f19])).
fof(f19, plain, mem(c_2Ebool_2EF, bool), file('/home/ubuntu/library/tptp/Problems/ITP/ITP006+2.p', mem_c_2Ebool_2EF)).
fof(f292, plain, ! [X0, X1] : (~ mem(X1, bool) | p(X1) | p(X0) | (X0 = X1) | ~ mem(X0, bool)), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ! [X0] : (! [X1] : ((X0 = X1) | ((~ p(X1) | ~ p(X0)) & (p(X1) | p(X0))) | ~ mem(X1, bool)) | ~ mem(X0, bool)), inference(nnf_transformation, [], [f109])).
fof(f109, plain, ! [X0] : (! [X1] : ((X0 = X1) | ~ (p(X0) <=> p(X1)) | ~ mem(X1, bool)) | ~ mem(X0, bool)), inference(flattening, [], [f108])).
fof(f108, plain, ! [X0] : (! [X1] : (((X0 = X1) | ~ (p(X0) <=> p(X1))) | ~ mem(X1, bool)) | ~ mem(X0, bool)), inference(ennf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : (mem(X0, bool) => ! [X1] : (mem(X1, bool) => ((p(X0) <=> p(X1)) => (X0 = X1)))), inference(rectify, [], [f5])).
fof(f5, plain, ! [X4] : (mem(X4, bool) => ! [X5] : (mem(X5, bool) => ((p(X4) <=> p(X5)) => (X4 = X5)))), file('/home/ubuntu/library/tptp/Problems/ITP/ITP006+2.p', boolext)).
fof(f1759, plain, (! [X0] : (~ p(ap(sK35, ap(sK34, X0))) | ~ mem(X0, sK33)) | ~ spl37_40), inference(avatar_component_clause, [], [f1758])).
fof(f797, plain, ! [X0] : (p(ap(sK35, X0)) | ~ mem(X0, sK32) | (c_2Ebool_2EF = ap(sK36, X0))), inference(duplicate_literal_removal, [], [f796])).
fof(f796, plain, ! [X0] : ((c_2Ebool_2EF = ap(sK36, X0)) | ~ mem(X0, sK32) | p(ap(sK35, X0)) | ~ mem(X0, sK32)), inference(resolution, [], [f712, f492])).
fof(f492, plain, ! [X5] : (~ p(ap(sK36, X5)) | p(ap(sK35, X5)) | ~ mem(X5, sK32)), inference(cnf_transformation, [], [f287])).
fof(f712, plain, ! [X1] : (p(ap(sK36, X1)) | (c_2Ebool_2EF = ap(sK36, X1)) | ~ mem(X1, sK32)), inference(resolution, [], [f677, f672])).
fof(f672, plain, ! [X12] : (mem(ap(sK36, X12), bool) | ~ mem(X12, sK32)), inference(resolution, [], [f291, f491])).
fof(f491, plain, mem(sK36, arr(sK32, bool)), inference(cnf_transformation, [], [f287])).
fof(f1767, plain, ~ spl37_39, inference(avatar_contradiction_clause, [], [f1761])).
fof(f1761, plain, ($false | ~ spl37_39), inference(resolution, [], [f1756, f489])).
fof(f1756, plain, (! [X2] : ~ mem(X2, arr(sK33, sK32)) | ~ spl37_39), inference(avatar_component_clause, [], [f1755])).
fof(f1755, plain, (spl37_39 <=> ! [X2] : ~ mem(X2, arr(sK33, sK32))), introduced(avatar_definition, [new_symbols(naming, [spl37_39])])).
fof(f1760, plain, (spl37_39 | spl37_34 | spl37_40), inference(avatar_split_clause, [], [f1753, f1758, f1703, f1755])).
fof(f1703, plain, (spl37_34 <=> ! [X31] : ~ mem(X31, arr(sK32, bool))), introduced(avatar_definition, [new_symbols(naming, [spl37_34])])).
fof(f1753, plain, ! [X2, X0, X1] : (~ mem(X0, sK33) | ~ p(ap(sK35, ap(sK34, X0))) | ~ mem(X1, arr(sK32, bool)) | ~ mem(X2, arr(sK33, sK32))), inference(subsumption_resolution, [], [f1752, f488])).
fof(f1752, plain, ! [X2, X0, X1] : (~ mem(X0, sK33) | ~ p(ap(sK35, ap(sK34, X0))) | ~ mem(X1, arr(sK32, bool)) | ~ mem(X2, arr(sK33, sK32)) | ~ ne(sK33)), inference(subsumption_resolution, [], [f1751, f487])).
fof(f487, plain, ne(sK32), inference(cnf_transformation, [], [f287])).
fof(f1751, plain, ! [X2, X0, X1] : (~ mem(X0, sK33) | ~ p(ap(sK35, ap(sK34, X0))) | ~ mem(X1, arr(sK32, bool)) | ~ mem(X2, arr(sK33, sK32)) | ~ ne(sK32) | ~ ne(sK33)), inference(subsumption_resolution, [], [f1750, f489])).
fof(f1750, plain, ! [X2, X0, X1] : (~ mem(X0, sK33) | ~ p(ap(sK35, ap(sK34, X0))) | ~ mem(sK34, arr(sK33, sK32)) | ~ mem(X1, arr(sK32, bool)) | ~ mem(X2, arr(sK33, sK32)) | ~ ne(sK32) | ~ ne(sK33)), inference(subsumption_resolution, [], [f1714, f490])).
fof(f1714, plain, ! [X2, X0, X1] : (~ mem(X0, sK33) | ~ p(ap(sK35, ap(sK34, X0))) | ~ mem(sK35, arr(sK32, bool)) | ~ mem(sK34, arr(sK33, sK32)) | ~ mem(X1, arr(sK32, bool)) | ~ mem(X2, arr(sK33, sK32)) | ~ ne(sK32) | ~ ne(sK33)), inference(resolution, [], [f393, f493])).
fof(f493, plain, p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(sK33, sK32), sK34), sK35)), inference(cnf_transformation, [], [f287])).
fof(f393, plain, ! [X4, X2, X0, X7, X5, X3, X1] : (~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X4), X5)) | ~ mem(X7, X0) | ~ p(ap(X5, ap(X4, X7))) | ~ mem(X5, arr(X1, bool)) | ~ mem(X4, arr(X0, X1)) | ~ mem(X3, arr(X1, bool)) | ~ mem(X2, arr(X0, X1)) | ~ ne(X1) | ~ ne(X0)), inference(cnf_transformation, [], [f241])).
fof(f241, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((sP4(X0, X1) & sP3(X0, X1) & ! [X4] : (! [X5] : (((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X4), X5)) | (p(ap(X5, ap(X4, sK31(X0, X4, X5)))) & mem(sK31(X0, X4, X5), X0))) & (! [X7] : (~ p(ap(X5, ap(X4, X7))) | ~ mem(X7, X0)) | ~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X4), X5)))) | ~ mem(X5, arr(X1, bool))) | ~ mem(X4, arr(X0, X1))) & sP5(X0, X1) & sP2(X2, X3, X0, X1) & sP1(X2, X3, X0, X1)) | ~ mem(X3, arr(X1, bool))) | ~ mem(X2, arr(X0, X1))) | ~ ne(X1)) | ~ ne(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK31])], [f239, f240])).
fof(f240, plain, ! [X5, X4, X0] : (? [X6] : (p(ap(X5, ap(X4, X6))) & mem(X6, X0)) => (p(ap(X5, ap(X4, sK31(X0, X4, X5)))) & mem(sK31(X0, X4, X5), X0))), introduced(choice_axiom, [])).
fof(f239, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((sP4(X0, X1) & sP3(X0, X1) & ! [X4] : (! [X5] : (((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X4), X5)) | ? [X6] : (p(ap(X5, ap(X4, X6))) & mem(X6, X0))) & (! [X7] : (~ p(ap(X5, ap(X4, X7))) | ~ mem(X7, X0)) | ~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X4), X5)))) | ~ mem(X5, arr(X1, bool))) | ~ mem(X4, arr(X0, X1))) & sP5(X0, X1) & sP2(X2, X3, X0, X1) & sP1(X2, X3, X0, X1)) | ~ mem(X3, arr(X1, bool))) | ~ mem(X2, arr(X0, X1))) | ~ ne(X1)) | ~ ne(X0)), inference(rectify, [], [f238])).
fof(f238, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((sP4(X0, X1) & sP3(X0, X1) & ! [X12] : (! [X13] : (((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X12), X13)) | ? [X14] : (p(ap(X13, ap(X12, X14))) & mem(X14, X0))) & (! [X14] : (~ p(ap(X13, ap(X12, X14))) | ~ mem(X14, X0)) | ~ p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X12), X13)))) | ~ mem(X13, arr(X1, bool))) | ~ mem(X12, arr(X0, X1))) & sP5(X0, X1) & sP2(X2, X3, X0, X1) & sP1(X2, X3, X0, X1)) | ~ mem(X3, arr(X1, bool))) | ~ mem(X2, arr(X0, X1))) | ~ ne(X1)) | ~ ne(X0)), inference(nnf_transformation, [], [f164])).
fof(f164, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((sP4(X0, X1) & sP3(X0, X1) & ! [X12] : (! [X13] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X12), X13)) <=> ! [X14] : (~ p(ap(X13, ap(X12, X14))) | ~ mem(X14, X0))) | ~ mem(X13, arr(X1, bool))) | ~ mem(X12, arr(X0, X1))) & sP5(X0, X1) & sP2(X2, X3, X0, X1) & sP1(X2, X3, X0, X1)) | ~ mem(X3, arr(X1, bool))) | ~ mem(X2, arr(X0, X1))) | ~ ne(X1)) | ~ ne(X0)), inference(definition_folding, [], [f139, e163, e162, e161, e160, e159])).
fof(f159, plain, ! [X2, X3, X0, X1] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS(X0, X1), X2), X3)) <=> ! [X20] : (? [X21] : (p(ap(X3, ap(X2, X21))) & mem(X21, X0)) | ~ p(ap(X3, X20)) | ~ mem(X20, X1))) | ~ sP1(X2, X3, X0, X1)), inference(usedef, [], [e159])).
fof(e159, plain, ! [X2, X3, X0, X1] : (sP1(X2, X3, X0, X1) <=> (p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS(X0, X1), X2), X3)) <=> ! [X20] : (? [X21] : (p(ap(X3, ap(X2, X21))) & mem(X21, X0)) | ~ p(ap(X3, X20)) | ~ mem(X20, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f160, plain, ! [X2, X3, X0, X1] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL(X0, X1), X2), X3)) <=> ! [X18] : (? [X19] : (~ p(ap(X3, ap(X2, X19))) & mem(X19, X0)) | p(ap(X3, X18)) | ~ mem(X18, X1))) | ~ sP2(X2, X3, X0, X1)), inference(usedef, [], [e160])).
fof(e160, plain, ! [X2, X3, X0, X1] : (sP2(X2, X3, X0, X1) <=> (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL(X0, X1), X2), X3)) <=> ! [X18] : (? [X19] : (~ p(ap(X3, ap(X2, X19))) & mem(X19, X0)) | p(ap(X3, X18)) | ~ mem(X18, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f161, plain, ! [X0, X1] : (! [X8] : (! [X9] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__GAP(X0, X1), X8), X9)) <=> ! [X10] : (? [X11] : ((ap(X8, X11) = X10) & mem(X11, X0)) | ~ p(ap(X9, X10)) | ~ mem(X10, X1))) | ~ mem(X9, arr(X1, bool))) | ~ mem(X8, arr(X0, X1))) | ~ sP3(X0, X1)), inference(usedef, [], [e161])).
fof(e161, plain, ! [X0, X1] : (sP3(X0, X1) <=> ! [X8] : (! [X9] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__GAP(X0, X1), X8), X9)) <=> ! [X10] : (? [X11] : ((ap(X8, X11) = X10) & mem(X11, X0)) | ~ p(ap(X9, X10)) | ~ mem(X10, X1))) | ~ mem(X9, arr(X1, bool))) | ~ mem(X8, arr(X0, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f162, plain, ! [X0, X1] : (! [X4] : (! [X5] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__GAP(X0, X1), X4), X5)) <=> ! [X6] : (? [X7] : ((ap(X4, X7) = X6) & mem(X7, X0)) | p(ap(X5, X6)) | ~ mem(X6, X1))) | ~ mem(X5, arr(X1, bool))) | ~ mem(X4, arr(X0, X1))) | ~ sP4(X0, X1)), inference(usedef, [], [e162])).
fof(e162, plain, ! [X0, X1] : (sP4(X0, X1) <=> ! [X4] : (! [X5] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__GAP(X0, X1), X4), X5)) <=> ! [X6] : (? [X7] : ((ap(X4, X7) = X6) & mem(X7, X0)) | p(ap(X5, X6)) | ~ mem(X6, X1))) | ~ mem(X5, arr(X1, bool))) | ~ mem(X4, arr(X0, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f163, plain, ! [X0, X1] : (! [X15] : (! [X16] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__POINT(X0, X1), X15), X16)) <=> ! [X17] : (p(ap(X16, ap(X15, X17))) | ~ mem(X17, X0))) | ~ mem(X16, arr(X1, bool))) | ~ mem(X15, arr(X0, X1))) | ~ sP5(X0, X1)), inference(usedef, [], [e163])).
fof(e163, plain, ! [X0, X1] : (sP5(X0, X1) <=> ! [X15] : (! [X16] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__POINT(X0, X1), X15), X16)) <=> ! [X17] : (p(ap(X16, ap(X15, X17))) | ~ mem(X17, X0))) | ~ mem(X16, arr(X1, bool))) | ~ mem(X15, arr(X0, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f139, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((! [X4] : (! [X5] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__GAP(X0, X1), X4), X5)) <=> ! [X6] : (? [X7] : ((ap(X4, X7) = X6) & mem(X7, X0)) | p(ap(X5, X6)) | ~ mem(X6, X1))) | ~ mem(X5, arr(X1, bool))) | ~ mem(X4, arr(X0, X1))) & ! [X8] : (! [X9] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__GAP(X0, X1), X8), X9)) <=> ! [X10] : (? [X11] : ((ap(X8, X11) = X10) & mem(X11, X0)) | ~ p(ap(X9, X10)) | ~ mem(X10, X1))) | ~ mem(X9, arr(X1, bool))) | ~ mem(X8, arr(X0, X1))) & ! [X12] : (! [X13] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X12), X13)) <=> ! [X14] : (~ p(ap(X13, ap(X12, X14))) | ~ mem(X14, X0))) | ~ mem(X13, arr(X1, bool))) | ~ mem(X12, arr(X0, X1))) & ! [X15] : (! [X16] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__POINT(X0, X1), X15), X16)) <=> ! [X17] : (p(ap(X16, ap(X15, X17))) | ~ mem(X17, X0))) | ~ mem(X16, arr(X1, bool))) | ~ mem(X15, arr(X0, X1))) & (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL(X0, X1), X2), X3)) <=> ! [X18] : (? [X19] : (~ p(ap(X3, ap(X2, X19))) & mem(X19, X0)) | p(ap(X3, X18)) | ~ mem(X18, X1))) & (p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS(X0, X1), X2), X3)) <=> ! [X20] : (? [X21] : (p(ap(X3, ap(X2, X21))) & mem(X21, X0)) | ~ p(ap(X3, X20)) | ~ mem(X20, X1)))) | ~ mem(X3, arr(X1, bool))) | ~ mem(X2, arr(X0, X1))) | ~ ne(X1)) | ~ ne(X0)), inference(flattening, [], [f138])).
fof(f138, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((! [X4] : (! [X5] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__GAP(X0, X1), X4), X5)) <=> ! [X6] : ((? [X7] : ((ap(X4, X7) = X6) & mem(X7, X0)) | p(ap(X5, X6))) | ~ mem(X6, X1))) | ~ mem(X5, arr(X1, bool))) | ~ mem(X4, arr(X0, X1))) & ! [X8] : (! [X9] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__GAP(X0, X1), X8), X9)) <=> ! [X10] : ((? [X11] : ((ap(X8, X11) = X10) & mem(X11, X0)) | ~ p(ap(X9, X10))) | ~ mem(X10, X1))) | ~ mem(X9, arr(X1, bool))) | ~ mem(X8, arr(X0, X1))) & ! [X12] : (! [X13] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X12), X13)) <=> ! [X14] : (~ p(ap(X13, ap(X12, X14))) | ~ mem(X14, X0))) | ~ mem(X13, arr(X1, bool))) | ~ mem(X12, arr(X0, X1))) & ! [X15] : (! [X16] : ((p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__POINT(X0, X1), X15), X16)) <=> ! [X17] : (p(ap(X16, ap(X15, X17))) | ~ mem(X17, X0))) | ~ mem(X16, arr(X1, bool))) | ~ mem(X15, arr(X0, X1))) & (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL(X0, X1), X2), X3)) <=> ! [X18] : ((? [X19] : (~ p(ap(X3, ap(X2, X19))) & mem(X19, X0)) | p(ap(X3, X18))) | ~ mem(X18, X1))) & (p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS(X0, X1), X2), X3)) <=> ! [X20] : ((? [X21] : (p(ap(X3, ap(X2, X21))) & mem(X21, X0)) | ~ p(ap(X3, X20))) | ~ mem(X20, X1)))) | ~ mem(X3, arr(X1, bool))) | ~ mem(X2, arr(X0, X1))) | ~ ne(X1)) | ~ ne(X0)), inference(ennf_transformation, [], [f85])).
fof(f85, plain, ! [X0] : (ne(X0) => ! [X1] : (ne(X1) => ! [X2] : (mem(X2, arr(X0, X1)) => ! [X3] : (mem(X3, arr(X1, bool)) => (! [X4] : (mem(X4, arr(X0, X1)) => ! [X5] : (mem(X5, arr(X1, bool)) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__GAP(X0, X1), X4), X5)) <=> ! [X6] : (mem(X6, X1) => (~ p(ap(X5, X6)) => ? [X7] : ((ap(X4, X7) = X6) & mem(X7, X0))))))) & ! [X8] : (mem(X8, arr(X0, X1)) => ! [X9] : (mem(X9, arr(X1, bool)) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__GAP(X0, X1), X8), X9)) <=> ! [X10] : (mem(X10, X1) => (p(ap(X9, X10)) => ? [X11] : ((ap(X8, X11) = X10) & mem(X11, X0))))))) & ! [X12] : (mem(X12, arr(X0, X1)) => ! [X13] : (mem(X13, arr(X1, bool)) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X12), X13)) <=> ! [X14] : (mem(X14, X0) => ~ p(ap(X13, ap(X12, X14))))))) & ! [X15] : (mem(X15, arr(X0, X1)) => ! [X16] : (mem(X16, arr(X1, bool)) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__POINT(X0, X1), X15), X16)) <=> ! [X17] : (mem(X17, X0) => p(ap(X16, ap(X15, X17))))))) & (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL(X0, X1), X2), X3)) <=> ! [X18] : (mem(X18, X1) => (~ p(ap(X3, X18)) => ? [X19] : (~ p(ap(X3, ap(X2, X19))) & mem(X19, X0))))) & (p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS(X0, X1), X2), X3)) <=> ! [X20] : (mem(X20, X1) => (p(ap(X3, X20)) => ? [X21] : (p(ap(X3, ap(X2, X21))) & mem(X21, X0)))))))))), inference(rectify, [], [f40])).
fof(f40, plain, ! [X8] : (ne(X8) => ! [X9] : (ne(X9) => ! [X19] : (mem(X19, arr(X8, X9)) => ! [X20] : (mem(X20, arr(X9, bool)) => (! [X35] : (mem(X35, arr(X8, X9)) => ! [X36] : (mem(X36, arr(X9, bool)) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__GAP(X8, X9), X35), X36)) <=> ! [X37] : (mem(X37, X9) => (~ p(ap(X36, X37)) => ? [X38] : ((ap(X35, X38) = X37) & mem(X38, X8))))))) & ! [X31] : (mem(X31, arr(X8, X9)) => ! [X32] : (mem(X32, arr(X9, bool)) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__GAP(X8, X9), X31), X32)) <=> ! [X33] : (mem(X33, X9) => (p(ap(X32, X33)) => ? [X34] : ((ap(X31, X34) = X33) & mem(X34, X8))))))) & ! [X28] : (mem(X28, arr(X8, X9)) => ! [X29] : (mem(X29, arr(X9, bool)) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X8, X9), X28), X29)) <=> ! [X30] : (mem(X30, X8) => ~ p(ap(X29, ap(X28, X30))))))) & ! [X25] : (mem(X25, arr(X8, X9)) => ! [X26] : (mem(X26, arr(X9, bool)) => (p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS__POINT(X8, X9), X25), X26)) <=> ! [X27] : (mem(X27, X8) => p(ap(X26, ap(X25, X27))))))) & (p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL(X8, X9), X19), X20)) <=> ! [X23] : (mem(X23, X9) => (~ p(ap(X20, X23)) => ? [X24] : (~ p(ap(X20, ap(X19, X24))) & mem(X24, X8))))) & (p(ap(ap(c_2EquantHeuristics_2EGUESS__EXISTS(X8, X9), X19), X20)) <=> ! [X21] : (mem(X21, X9) => (p(ap(X20, X21)) => ? [X22] : (p(ap(X20, ap(X19, X22))) & mem(X22, X8)))))))))), file('/home/ubuntu/library/tptp/Problems/ITP/ITP006+2.p', conj_thm_2EquantHeuristics_2EGUESS__REWRITES)).
fof(f1744, plain, (spl37_34 | spl37_37), inference(avatar_split_clause, [], [f1740, f1742, f1703])).
fof(f1740, plain, ! [X35, X33, X36, X34] : (p(ap(sK36, ap(X33, sK31(X34, X33, sK36)))) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X34, sK32), X33), sK36)) | ~ mem(X33, arr(X34, sK32)) | ~ mem(X35, arr(sK32, bool)) | ~ mem(X36, arr(X34, sK32)) | ~ ne(X34)), inference(subsumption_resolution, [], [f1725, f487])).
fof(f1725, plain, ! [X35, X33, X36, X34] : (p(ap(sK36, ap(X33, sK31(X34, X33, sK36)))) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X34, sK32), X33), sK36)) | ~ mem(X33, arr(X34, sK32)) | ~ mem(X35, arr(sK32, bool)) | ~ mem(X36, arr(X34, sK32)) | ~ ne(sK32) | ~ ne(X34)), inference(resolution, [], [f395, f491])).
fof(f395, plain, ! [X4, X2, X0, X5, X3, X1] : (~ mem(X5, arr(X1, bool)) | p(ap(X5, ap(X4, sK31(X0, X4, X5)))) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X4), X5)) | ~ mem(X4, arr(X0, X1)) | ~ mem(X3, arr(X1, bool)) | ~ mem(X2, arr(X0, X1)) | ~ ne(X1) | ~ ne(X0)), inference(cnf_transformation, [], [f241])).
fof(f1739, plain, ~ spl37_34, inference(avatar_contradiction_clause, [], [f1731])).
fof(f1731, plain, ($false | ~ spl37_34), inference(resolution, [], [f1704, f490])).
fof(f1704, plain, (! [X31] : ~ mem(X31, arr(sK32, bool)) | ~ spl37_34), inference(avatar_component_clause, [], [f1703])).
fof(f1713, plain, (spl37_34 | spl37_36), inference(avatar_split_clause, [], [f1709, f1711, f1703])).
fof(f1709, plain, ! [X35, X33, X36, X34] : (mem(sK31(X33, X34, sK36), X33) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X33, sK32), X34), sK36)) | ~ mem(X34, arr(X33, sK32)) | ~ mem(X35, arr(sK32, bool)) | ~ mem(X36, arr(X33, sK32)) | ~ ne(X33)), inference(subsumption_resolution, [], [f1672, f487])).
fof(f1672, plain, ! [X35, X33, X36, X34] : (mem(sK31(X33, X34, sK36), X33) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X33, sK32), X34), sK36)) | ~ mem(X34, arr(X33, sK32)) | ~ mem(X35, arr(sK32, bool)) | ~ mem(X36, arr(X33, sK32)) | ~ ne(sK32) | ~ ne(X33)), inference(resolution, [], [f394, f491])).
fof(f394, plain, ! [X4, X2, X0, X5, X3, X1] : (~ mem(X5, arr(X1, bool)) | mem(sK31(X0, X4, X5), X0) | p(ap(ap(c_2EquantHeuristics_2EGUESS__FORALL__POINT(X0, X1), X4), X5)) | ~ mem(X4, arr(X0, X1)) | ~ mem(X3, arr(X1, bool)) | ~ mem(X2, arr(X0, X1)) | ~ ne(X1) | ~ ne(X0)), inference(cnf_transformation, [], [f241])).