fof(f398, plain, $false, inference(subsumption_resolution, [], [f397, f226])).
fof(f226, plain, aSet0(xP), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((aElementOf0(X0, xP) | (~ aUpperBoundOfIn0(X0, xT, xU) & (~ sdtlseqdt0(sK19(X0), X0) & aElementOf0(sK19(X0), xT))) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xU)) & ((aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (sdtlseqdt0(X2, X0) | ~ aElementOf0(X2, xT)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) | ~ aElementOf0(X0, xP))) & aSet0(xP)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK19])], [f68, f132])).
fof(f132, plain, ! [X0] : (? [X1] : (~ sdtlseqdt0(X1, X0) & aElementOf0(X1, xT)) => (~ sdtlseqdt0(sK19(X0), X0) & aElementOf0(sK19(X0), xT))), introduced(choice_axiom, [])).
fof(f68, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((aElementOf0(X0, xP) | (~ aUpperBoundOfIn0(X0, xT, xU) & ? [X1] : (~ sdtlseqdt0(X1, X0) & aElementOf0(X1, xT))) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xU)) & ((aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (sdtlseqdt0(X2, X0) | ~ aElementOf0(X2, xT)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) | ~ aElementOf0(X0, xP))) & aSet0(xP)), inference(flattening, [], [f67])).
fof(f67, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((aElementOf0(X0, xP) | ((~ aUpperBoundOfIn0(X0, xT, xU) & ? [X1] : (~ sdtlseqdt0(X1, X0) & aElementOf0(X1, xT))) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xU))) & ((aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (sdtlseqdt0(X2, X0) | ~ aElementOf0(X2, xT)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) | ~ aElementOf0(X0, xP))) & aSet0(xP)), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((((aUpperBoundOfIn0(X0, xT, xU) | ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0))) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) => aElementOf0(X0, xP)) & (aElementOf0(X0, xP) => (aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (aElementOf0(X2, xT) => sdtlseqdt0(X2, X0)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)))) & aSet0(xP)), inference(rectify, [], [f27])).
fof(f27, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((((aUpperBoundOfIn0(X0, xT, xU) | ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0))) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) => aElementOf0(X0, xP)) & (aElementOf0(X0, xP) => (aUpperBoundOfIn0(X0, xT, xU) & ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)))) & aSet0(xP)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT385+4.p', m__1244)).
fof(f397, plain, ~ aSet0(xP), inference(subsumption_resolution, [], [f394, f288])).
fof(f288, plain, ~ sP2(xP), inference(resolution, [], [f192, f242])).
fof(f242, plain, ! [X0] : ~ aInfimumOfIn0(X0, xP, xU), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ! [X0] : (~ aInfimumOfIn0(X0, xP, xU) & (sP3(X0) | (~ aLowerBoundOfIn0(X0, xP, xU) & ((~ sdtlseqdt0(X0, sK21(X0)) & aElementOf0(sK21(X0), xP)) | ~ aElementOf0(X0, xU))) | ~ aElementOf0(X0, xU))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21])], [f137, f138])).
fof(f138, plain, ! [X0] : (? [X1] : (~ sdtlseqdt0(X0, X1) & aElementOf0(X1, xP)) => (~ sdtlseqdt0(X0, sK21(X0)) & aElementOf0(sK21(X0), xP))), introduced(choice_axiom, [])).
fof(f137, plain, ! [X0] : (~ aInfimumOfIn0(X0, xP, xU) & (sP3(X0) | (~ aLowerBoundOfIn0(X0, xP, xU) & (? [X1] : (~ sdtlseqdt0(X0, X1) & aElementOf0(X1, xP)) | ~ aElementOf0(X0, xU))) | ~ aElementOf0(X0, xU))), inference(rectify, [], [f76])).
fof(f76, plain, ! [X0] : (~ aInfimumOfIn0(X0, xP, xU) & (sP3(X0) | (~ aLowerBoundOfIn0(X0, xP, xU) & (? [X3] : (~ sdtlseqdt0(X0, X3) & aElementOf0(X3, xP)) | ~ aElementOf0(X0, xU))) | ~ aElementOf0(X0, xU))), inference(definition_folding, [], [f70, e75])).
fof(f75, plain, ! [X0] : (? [X1] : (~ sdtlseqdt0(X1, X0) & aLowerBoundOfIn0(X1, xP, xU) & ! [X2] : (sdtlseqdt0(X1, X2) | ~ aElementOf0(X2, xP)) & aElementOf0(X1, xU)) | ~ sP3(X0)), inference(usedef, [], [e75])).
fof(e75, plain, ! [X0] : (sP3(X0) <=> ? [X1] : (~ sdtlseqdt0(X1, X0) & aLowerBoundOfIn0(X1, xP, xU) & ! [X2] : (sdtlseqdt0(X1, X2) | ~ aElementOf0(X2, xP)) & aElementOf0(X1, xU))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f70, plain, ! [X0] : (~ aInfimumOfIn0(X0, xP, xU) & (? [X1] : (~ sdtlseqdt0(X1, X0) & aLowerBoundOfIn0(X1, xP, xU) & ! [X2] : (sdtlseqdt0(X1, X2) | ~ aElementOf0(X2, xP)) & aElementOf0(X1, xU)) | (~ aLowerBoundOfIn0(X0, xP, xU) & (? [X3] : (~ sdtlseqdt0(X0, X3) & aElementOf0(X3, xP)) | ~ aElementOf0(X0, xU))) | ~ aElementOf0(X0, xU))), inference(flattening, [], [f69])).
fof(f69, plain, ! [X0] : (~ aInfimumOfIn0(X0, xP, xU) & (? [X1] : (~ sdtlseqdt0(X1, X0) & (aLowerBoundOfIn0(X1, xP, xU) & ! [X2] : (sdtlseqdt0(X1, X2) | ~ aElementOf0(X2, xP)) & aElementOf0(X1, xU))) | (~ aLowerBoundOfIn0(X0, xP, xU) & (? [X3] : (~ sdtlseqdt0(X0, X3) & aElementOf0(X3, xP)) | ~ aElementOf0(X0, xU))) | ~ aElementOf0(X0, xU))), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ~ ? [X0] : (aInfimumOfIn0(X0, xP, xU) | (! [X1] : ((aLowerBoundOfIn0(X1, xP, xU) & ! [X2] : (aElementOf0(X2, xP) => sdtlseqdt0(X1, X2)) & aElementOf0(X1, xU)) => sdtlseqdt0(X1, X0)) & (aLowerBoundOfIn0(X0, xP, xU) | (! [X3] : (aElementOf0(X3, xP) => sdtlseqdt0(X0, X3)) & aElementOf0(X0, xU))) & aElementOf0(X0, xU))), inference(rectify, [], [f29])).
fof(f29, plain, ~ ? [X0] : (aInfimumOfIn0(X0, xP, xU) | (! [X1] : ((aLowerBoundOfIn0(X1, xP, xU) & ! [X2] : (aElementOf0(X2, xP) => sdtlseqdt0(X1, X2)) & aElementOf0(X1, xU)) => sdtlseqdt0(X1, X0)) & (aLowerBoundOfIn0(X0, xP, xU) | (! [X1] : (aElementOf0(X1, xP) => sdtlseqdt0(X0, X1)) & aElementOf0(X0, xU))) & aElementOf0(X0, xU))), inference(negated_conjecture, [], [f28])).
fof(f28, plain, ~ ? [X0] : (aInfimumOfIn0(X0, xP, xU) | (! [X1] : ((aLowerBoundOfIn0(X1, xP, xU) & ! [X2] : (aElementOf0(X2, xP) => sdtlseqdt0(X1, X2)) & aElementOf0(X1, xU)) => sdtlseqdt0(X1, X0)) & (aLowerBoundOfIn0(X0, xP, xU) | (! [X1] : (aElementOf0(X1, xP) => sdtlseqdt0(X0, X1)) & aElementOf0(X0, xU))) & aElementOf0(X0, xU))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT385+4.p', m__)).
fof(f192, plain, ! [X0] : (aInfimumOfIn0(sK14(X0), X0, xU) | ~ sP2(X0)), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ! [X0] : ((sP0(X0) & aInfimumOfIn0(sK14(X0), X0, xU) & sP1(sK14(X0), X0) & aLowerBoundOfIn0(sK14(X0), X0, xU) & ! [X2] : (sdtlseqdt0(sK14(X0), X2) | ~ aElementOf0(X2, X0)) & aElementOf0(sK14(X0), xU) & aElementOf0(sK14(X0), xU)) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14])], [f118, f119])).
fof(f119, plain, ! [X0] : (? [X1] : (sP0(X0) & aInfimumOfIn0(X1, X0, xU) & sP1(X1, X0) & aLowerBoundOfIn0(X1, X0, xU) & ! [X2] : (sdtlseqdt0(X1, X2) | ~ aElementOf0(X2, X0)) & aElementOf0(X1, xU) & aElementOf0(X1, xU)) => (sP0(X0) & aInfimumOfIn0(sK14(X0), X0, xU) & sP1(sK14(X0), X0) & aLowerBoundOfIn0(sK14(X0), X0, xU) & ! [X2] : (sdtlseqdt0(sK14(X0), X2) | ~ aElementOf0(X2, X0)) & aElementOf0(sK14(X0), xU) & aElementOf0(sK14(X0), xU))), introduced(choice_axiom, [])).
fof(f118, plain, ! [X0] : (? [X1] : (sP0(X0) & aInfimumOfIn0(X1, X0, xU) & sP1(X1, X0) & aLowerBoundOfIn0(X1, X0, xU) & ! [X2] : (sdtlseqdt0(X1, X2) | ~ aElementOf0(X2, X0)) & aElementOf0(X1, xU) & aElementOf0(X1, xU)) | ~ sP2(X0)), inference(rectify, [], [f117])).
fof(f117, plain, ! [X2] : (? [X4] : (sP0(X2) & aInfimumOfIn0(X4, X2, xU) & sP1(X4, X2) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU)) | ~ sP2(X2)), inference(nnf_transformation, [], [f73])).
fof(f73, plain, ! [X2] : (? [X4] : (sP0(X2) & aInfimumOfIn0(X4, X2, xU) & sP1(X4, X2) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU)) | ~ sP2(X2)), inference(usedef, [], [e73])).
fof(e73, plain, ! [X2] : (sP2(X2) <=> ? [X4] : (sP0(X2) & aInfimumOfIn0(X4, X2, xU) & sP1(X4, X2) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f394, plain, (sP2(xP) | ~ aSet0(xP)), inference(resolution, [], [f324, f207])).
fof(f207, plain, ! [X2] : (~ aElementOf0(sK18(X2), xU) | sP2(X2) | ~ aSet0(X2)), inference(cnf_transformation, [], [f131])).
fof(f131, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (sP2(X2) | (~ aSubsetOf0(X2, xU) & ((~ aElementOf0(sK18(X2), xU) & aElementOf0(sK18(X2), X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK18])], [f74, f130])).
fof(f130, plain, ! [X2] : (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) => (~ aElementOf0(sK18(X2), xU) & aElementOf0(sK18(X2), X2))), introduced(choice_axiom, [])).
fof(f74, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (sP2(X2) | (~ aSubsetOf0(X2, xU) & (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(definition_folding, [], [f64, e73, e72, e71])).
fof(f71, plain, ! [X2] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) | ~ sP0(X2)), inference(usedef, [], [e71])).
fof(e71, plain, ! [X2] : (sP0(X2) <=> ? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f72, plain, ! [X4, X2] : (! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU)))) | ~ sP1(X4, X2)), inference(usedef, [], [e72])).
fof(e72, plain, ! [X4, X2] : (sP1(X4, X2) <=> ! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f64, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (? [X4] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) & aInfimumOfIn0(X4, X2, xU) & ! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU)))) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU)) | (~ aSubsetOf0(X2, xU) & (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(flattening, [], [f63])).
fof(f63, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : ((sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1)) | (~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (? [X4] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) & aInfimumOfIn0(X4, X2, xU) & ! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU)))) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU)) | (~ aSubsetOf0(X2, xU) & (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(ennf_transformation, [], [f34])).
fof(f34, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : ((aElementOf0(X1, szDzozmdt0(xf)) & aElementOf0(X0, szDzozmdt0(xf))) => (sdtlseqdt0(X0, X1) => sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : ((aSubsetOf0(X2, xU) | (! [X3] : (aElementOf0(X3, X2) => aElementOf0(X3, xU)) & aSet0(X2))) => ? [X4] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : ((aUpperBoundOfIn0(X6, X2, xU) | (! [X7] : (aElementOf0(X7, X2) => sdtlseqdt0(X7, X6)) & aElementOf0(X6, xU))) => sdtlseqdt0(X5, X6)) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (aElementOf0(X8, X2) => sdtlseqdt0(X8, X5)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) & aInfimumOfIn0(X4, X2, xU) & ! [X9] : ((aLowerBoundOfIn0(X9, X2, xU) | (! [X10] : (aElementOf0(X10, X2) => sdtlseqdt0(X9, X10)) & aElementOf0(X9, xU))) => sdtlseqdt0(X9, X4)) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (aElementOf0(X11, X2) => sdtlseqdt0(X4, X11)) & aElementOf0(X4, xU) & aElementOf0(X4, xU))) & aSet0(xU)), inference(rectify, [], [f24])).
fof(f24, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : ((aElementOf0(X1, szDzozmdt0(xf)) & aElementOf0(X0, szDzozmdt0(xf))) => (sdtlseqdt0(X0, X1) => sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X0] : ((aSubsetOf0(X0, xU) | (! [X1] : (aElementOf0(X1, X0) => aElementOf0(X1, xU)) & aSet0(X0))) => ? [X1] : (? [X2] : (aSupremumOfIn0(X2, X0, xU) & ! [X3] : ((aUpperBoundOfIn0(X3, X0, xU) | (! [X4] : (aElementOf0(X4, X0) => sdtlseqdt0(X4, X3)) & aElementOf0(X3, xU))) => sdtlseqdt0(X2, X3)) & aUpperBoundOfIn0(X2, X0, xU) & ! [X3] : (aElementOf0(X3, X0) => sdtlseqdt0(X3, X2)) & aElementOf0(X2, xU) & aElementOf0(X2, xU)) & aInfimumOfIn0(X1, X0, xU) & ! [X2] : ((aLowerBoundOfIn0(X2, X0, xU) | (! [X3] : (aElementOf0(X3, X0) => sdtlseqdt0(X2, X3)) & aElementOf0(X2, xU))) => sdtlseqdt0(X2, X1)) & aLowerBoundOfIn0(X1, X0, xU) & ! [X2] : (aElementOf0(X2, X0) => sdtlseqdt0(X1, X2)) & aElementOf0(X1, xU) & aElementOf0(X1, xU))) & aSet0(xU)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT385+4.p', m__1123)).
fof(f324, plain, aElementOf0(sK18(xP), xU), inference(subsumption_resolution, [], [f323, f226])).
fof(f323, plain, (~ aSet0(xP) | aElementOf0(sK18(xP), xU)), inference(subsumption_resolution, [], [f294, f288])).
fof(f294, plain, (sP2(xP) | ~ aSet0(xP) | aElementOf0(sK18(xP), xU)), inference(resolution, [], [f206, f227])).
fof(f227, plain, ! [X0] : (~ aElementOf0(X0, xP) | aElementOf0(X0, xU)), inference(cnf_transformation, [], [f133])).
fof(f206, plain, ! [X2] : (aElementOf0(sK18(X2), X2) | sP2(X2) | ~ aSet0(X2)), inference(cnf_transformation, [], [f131])).