fof(f6799, plain, $false, inference(avatar_sat_refutation, [], [f1599, f1630, f3329, f3342, f3444, f3565, f3658, f6311, f6680, f6710, f6721, f6766, f6785, f6798])).
fof(f6798, plain, (~ spl33_99 | ~ spl33_561), inference(avatar_contradiction_clause, [], [f6797])).
fof(f6797, plain, ($false | (~ spl33_99 | ~ spl33_561)), inference(subsumption_resolution, [], [f6789, f1629])).
fof(f1629, plain, (sP0(k7_conlat_2(sK5)) | ~ spl33_99), inference(avatar_component_clause, [], [f1627])).
fof(f1627, plain, (spl33_99 <=> sP0(k7_conlat_2(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl33_99])])).
fof(f6789, plain, (~ sP0(k7_conlat_2(sK5)) | ~ spl33_561), inference(resolution, [], [f6306, f369])).
fof(f369, plain, ! [X0] : (~ v3_struct_0(k11_conlat_1(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f255])).
fof(f255, plain, ! [X0] : ((v4_lattice3(k11_conlat_1(X0)) & v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ sP0(X0)), inference(nnf_transformation, [], [f219])).
fof(f219, plain, ! [X0] : ((v4_lattice3(k11_conlat_1(X0)) & v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ sP0(X0)), inference(usedef, [], [e219])).
fof(e219, plain, ! [X0] : (sP0(X0) <=> (v4_lattice3(k11_conlat_1(X0)) & v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f6306, plain, (v3_struct_0(k11_conlat_1(k7_conlat_2(sK5))) | ~ spl33_561), inference(avatar_component_clause, [], [f6304])).
fof(f6304, plain, (spl33_561 <=> v3_struct_0(k11_conlat_1(k7_conlat_2(sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl33_561])])).
fof(f6785, plain, (~ spl33_99 | spl33_562), inference(avatar_contradiction_clause, [], [f6784])).
fof(f6784, plain, ($false | (~ spl33_99 | spl33_562)), inference(subsumption_resolution, [], [f6779, f1629])).
fof(f6779, plain, (~ sP0(k7_conlat_2(sK5)) | spl33_562), inference(resolution, [], [f6310, f377])).
fof(f377, plain, ! [X0] : (v10_lattices(k11_conlat_1(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f255])).
fof(f6310, plain, (~ v10_lattices(k11_conlat_1(k7_conlat_2(sK5))) | spl33_562), inference(avatar_component_clause, [], [f6308])).
fof(f6308, plain, (spl33_562 <=> v10_lattices(k11_conlat_1(k7_conlat_2(sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl33_562])])).
fof(f6766, plain, (~ spl33_223 | ~ spl33_558), inference(avatar_contradiction_clause, [], [f6765])).
fof(f6765, plain, ($false | (~ spl33_223 | ~ spl33_558)), inference(subsumption_resolution, [], [f6759, f3705])).
fof(f3705, plain, (sP1(k11_conlat_1(sK5)) | ~ spl33_223), inference(subsumption_resolution, [], [f3696, f604])).
fof(f604, plain, sP3(sK5), inference(subsumption_resolution, [], [f600, f293])).
fof(f293, plain, ~ v3_conlat_1(sK5), inference(cnf_transformation, [], [f230])).
fof(f230, plain, (~ r1_filter_1(k11_conlat_1(k7_conlat_2(sK5)), k1_lattice2(k11_conlat_1(sK5))) & l2_conlat_1(sK5) & ~ v3_conlat_1(sK5)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5])], [f111, f229])).
fof(f229, plain, (? [X0] : (~ r1_filter_1(k11_conlat_1(k7_conlat_2(X0)), k1_lattice2(k11_conlat_1(X0))) & l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (~ r1_filter_1(k11_conlat_1(k7_conlat_2(sK5)), k1_lattice2(k11_conlat_1(sK5))) & l2_conlat_1(sK5) & ~ v3_conlat_1(sK5))), introduced(choice_axiom, [])).
fof(f111, plain, ? [X0] : (~ r1_filter_1(k11_conlat_1(k7_conlat_2(X0)), k1_lattice2(k11_conlat_1(X0))) & l2_conlat_1(X0) & ~ v3_conlat_1(X0)), inference(flattening, [], [f110])).
fof(f110, plain, ? [X0] : (~ r1_filter_1(k11_conlat_1(k7_conlat_2(X0)), k1_lattice2(k11_conlat_1(X0))) & (l2_conlat_1(X0) & ~ v3_conlat_1(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => r1_filter_1(k11_conlat_1(k7_conlat_2(X0)), k1_lattice2(k11_conlat_1(X0)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => r1_filter_1(k11_conlat_1(k7_conlat_2(X0)), k1_lattice2(k11_conlat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', t22_conlat_2)).
fof(f600, plain, (sP3(sK5) | v3_conlat_1(sK5)), inference(resolution, [], [f435, f294])).
fof(f294, plain, l2_conlat_1(sK5), inference(cnf_transformation, [], [f230])).
fof(f435, plain, ! [X0] : (~ l2_conlat_1(X0) | sP3(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f226])).
fof(f226, plain, ! [X0] : (sP3(X0) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(definition_folding, [], [f185, e225])).
fof(f225, plain, ! [X0] : ((v4_lattice3(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ sP3(X0)), inference(usedef, [], [e225])).
fof(e225, plain, ! [X0] : (sP3(X0) <=> (v4_lattice3(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f185, plain, ! [X0] : ((v4_lattice3(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f184])).
fof(f184, plain, ! [X0] : ((v4_lattice3(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (v4_lattice3(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', fc5_conlat_1)).
fof(f3696, plain, (sP1(k11_conlat_1(sK5)) | ~ sP3(sK5) | ~ spl33_223), inference(resolution, [], [f3443, f755])).
fof(f755, plain, ! [X0] : (~ l3_lattices(k11_conlat_1(X0)) | sP1(k11_conlat_1(X0)) | ~ sP3(X0)), inference(subsumption_resolution, [], [f754, f425])).
fof(f425, plain, ! [X0] : (~ v3_struct_0(k11_conlat_1(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f258])).
fof(f258, plain, ! [X0] : ((v4_lattice3(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ sP3(X0)), inference(nnf_transformation, [], [f225])).
fof(f754, plain, ! [X0] : (~ l3_lattices(k11_conlat_1(X0)) | sP1(k11_conlat_1(X0)) | v3_struct_0(k11_conlat_1(X0)) | ~ sP3(X0)), inference(subsumption_resolution, [], [f751, f433])).
fof(f433, plain, ! [X0] : (v10_lattices(k11_conlat_1(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f258])).
fof(f751, plain, ! [X0] : (~ l3_lattices(k11_conlat_1(X0)) | sP1(k11_conlat_1(X0)) | ~ v10_lattices(k11_conlat_1(X0)) | v3_struct_0(k11_conlat_1(X0)) | ~ sP3(X0)), inference(resolution, [], [f399, f434])).
fof(f434, plain, ! [X0] : (v4_lattice3(k11_conlat_1(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f258])).
fof(f399, plain, ! [X0] : (~ v4_lattice3(X0) | ~ l3_lattices(X0) | sP1(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f222])).
fof(f222, plain, ! [X0] : (sP1(X0) | ~ l3_lattices(X0) | ~ v4_lattice3(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(definition_folding, [], [f171, e221])).
fof(f221, plain, ! [X0] : ((v4_lattice3(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP1(X0)), inference(usedef, [], [e221])).
fof(e221, plain, ! [X0] : (sP1(X0) <=> (v4_lattice3(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f171, plain, ! [X0] : ((v4_lattice3(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v4_lattice3(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f170])).
fof(f170, plain, ! [X0] : ((v4_lattice3(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v4_lattice3(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f60])).
fof(f60, plain, ! [X0] : ((l3_lattices(X0) & v4_lattice3(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v4_lattice3(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', fc2_conlat_2)).
fof(f3443, plain, (l3_lattices(k11_conlat_1(sK5)) | ~ spl33_223), inference(avatar_component_clause, [], [f3441])).
fof(f3441, plain, (spl33_223 <=> l3_lattices(k11_conlat_1(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl33_223])])).
fof(f6759, plain, (~ sP1(k11_conlat_1(sK5)) | ~ spl33_558), inference(resolution, [], [f6294, f389])).
fof(f389, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f256])).
fof(f256, plain, ! [X0] : ((v4_lattice3(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP1(X0)), inference(nnf_transformation, [], [f221])).
fof(f6294, plain, (v3_struct_0(k1_lattice2(k11_conlat_1(sK5))) | ~ spl33_558), inference(avatar_component_clause, [], [f6292])).
fof(f6292, plain, (spl33_558 <=> v3_struct_0(k1_lattice2(k11_conlat_1(sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl33_558])])).
fof(f6721, plain, (~ spl33_223 | spl33_560), inference(avatar_contradiction_clause, [], [f6720])).
fof(f6720, plain, ($false | (~ spl33_223 | spl33_560)), inference(subsumption_resolution, [], [f6719, f3443])).
fof(f6719, plain, (~ l3_lattices(k11_conlat_1(sK5)) | spl33_560), inference(resolution, [], [f6302, f336])).
fof(f336, plain, ! [X0] : (l3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ! [X0] : ((l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : (l3_lattices(X0) => (l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_k1_lattice2)).
fof(f6302, plain, (~ l3_lattices(k1_lattice2(k11_conlat_1(sK5))) | spl33_560), inference(avatar_component_clause, [], [f6300])).
fof(f6300, plain, (spl33_560 <=> l3_lattices(k1_lattice2(k11_conlat_1(sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl33_560])])).
fof(f6710, plain, (~ spl33_223 | spl33_559), inference(avatar_contradiction_clause, [], [f6709])).
fof(f6709, plain, ($false | (~ spl33_223 | spl33_559)), inference(subsumption_resolution, [], [f6706, f3705])).
fof(f6706, plain, (~ sP1(k11_conlat_1(sK5)) | spl33_559), inference(resolution, [], [f6298, f397])).
fof(f397, plain, ! [X0] : (v10_lattices(k1_lattice2(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f256])).
fof(f6298, plain, (~ v10_lattices(k1_lattice2(k11_conlat_1(sK5))) | spl33_559), inference(avatar_component_clause, [], [f6296])).
fof(f6296, plain, (spl33_559 <=> v10_lattices(k1_lattice2(k11_conlat_1(sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl33_559])])).
fof(f6680, plain, (~ spl33_92 | spl33_96 | spl33_557), inference(avatar_contradiction_clause, [], [f6679])).
fof(f6679, plain, ($false | (~ spl33_92 | spl33_96 | spl33_557)), inference(subsumption_resolution, [], [f6678, f1614])).
fof(f1614, plain, (~ v3_conlat_1(k7_conlat_2(sK5)) | spl33_96), inference(avatar_component_clause, [], [f1613])).
fof(f1613, plain, (spl33_96 <=> v3_conlat_1(k7_conlat_2(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl33_96])])).
fof(f6678, plain, (v3_conlat_1(k7_conlat_2(sK5)) | (~ spl33_92 | spl33_557)), inference(subsumption_resolution, [], [f6677, f1598])).
fof(f1598, plain, (l2_conlat_1(k7_conlat_2(sK5)) | ~ spl33_92), inference(avatar_component_clause, [], [f1596])).
fof(f1596, plain, (spl33_92 <=> l2_conlat_1(k7_conlat_2(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl33_92])])).
fof(f6677, plain, (~ l2_conlat_1(k7_conlat_2(sK5)) | v3_conlat_1(k7_conlat_2(sK5)) | spl33_557), inference(resolution, [], [f6290, f334])).
fof(f334, plain, ! [X0] : (l3_lattices(k11_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ! [X0] : ((l3_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f141])).
fof(f141, plain, ! [X0] : ((l3_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (l3_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_k11_conlat_1)).
fof(f6290, plain, (~ l3_lattices(k11_conlat_1(k7_conlat_2(sK5))) | spl33_557), inference(avatar_component_clause, [], [f6288])).
fof(f6288, plain, (spl33_557 <=> l3_lattices(k11_conlat_1(k7_conlat_2(sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl33_557])])).
fof(f6311, plain, (~ spl33_557 | spl33_558 | ~ spl33_559 | ~ spl33_560 | spl33_561 | ~ spl33_562), inference(avatar_split_clause, [], [f6286, f6308, f6304, f6300, f6296, f6292, f6288])).
fof(f6286, plain, (~ v10_lattices(k11_conlat_1(k7_conlat_2(sK5))) | v3_struct_0(k11_conlat_1(k7_conlat_2(sK5))) | ~ l3_lattices(k1_lattice2(k11_conlat_1(sK5))) | ~ v10_lattices(k1_lattice2(k11_conlat_1(sK5))) | v3_struct_0(k1_lattice2(k11_conlat_1(sK5))) | ~ l3_lattices(k11_conlat_1(k7_conlat_2(sK5)))), inference(subsumption_resolution, [], [f6285, f293])).
fof(f6285, plain, (~ v10_lattices(k11_conlat_1(k7_conlat_2(sK5))) | v3_struct_0(k11_conlat_1(k7_conlat_2(sK5))) | ~ l3_lattices(k1_lattice2(k11_conlat_1(sK5))) | ~ v10_lattices(k1_lattice2(k11_conlat_1(sK5))) | v3_struct_0(k1_lattice2(k11_conlat_1(sK5))) | v3_conlat_1(sK5) | ~ l3_lattices(k11_conlat_1(k7_conlat_2(sK5)))), inference(subsumption_resolution, [], [f6282, f294])).
fof(f6282, plain, (~ v10_lattices(k11_conlat_1(k7_conlat_2(sK5))) | v3_struct_0(k11_conlat_1(k7_conlat_2(sK5))) | ~ l3_lattices(k1_lattice2(k11_conlat_1(sK5))) | ~ v10_lattices(k1_lattice2(k11_conlat_1(sK5))) | v3_struct_0(k1_lattice2(k11_conlat_1(sK5))) | ~ l2_conlat_1(sK5) | v3_conlat_1(sK5) | ~ l3_lattices(k11_conlat_1(k7_conlat_2(sK5)))), inference(resolution, [], [f1903, f295])).
fof(f295, plain, ~ r1_filter_1(k11_conlat_1(k7_conlat_2(sK5)), k1_lattice2(k11_conlat_1(sK5))), inference(cnf_transformation, [], [f230])).
fof(f1903, plain, ! [X0] : (r1_filter_1(k11_conlat_1(k7_conlat_2(X0)), k1_lattice2(k11_conlat_1(X0))) | ~ v10_lattices(k11_conlat_1(k7_conlat_2(X0))) | v3_struct_0(k11_conlat_1(k7_conlat_2(X0))) | ~ l3_lattices(k1_lattice2(k11_conlat_1(X0))) | ~ v10_lattices(k1_lattice2(k11_conlat_1(X0))) | v3_struct_0(k1_lattice2(k11_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0) | ~ l3_lattices(k11_conlat_1(k7_conlat_2(X0)))), inference(duplicate_literal_removal, [], [f1902])).
fof(f1902, plain, ! [X0] : (~ l3_lattices(k11_conlat_1(k7_conlat_2(X0))) | ~ v10_lattices(k11_conlat_1(k7_conlat_2(X0))) | v3_struct_0(k11_conlat_1(k7_conlat_2(X0))) | ~ l3_lattices(k1_lattice2(k11_conlat_1(X0))) | ~ v10_lattices(k1_lattice2(k11_conlat_1(X0))) | v3_struct_0(k1_lattice2(k11_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0) | r1_filter_1(k11_conlat_1(k7_conlat_2(X0)), k1_lattice2(k11_conlat_1(X0))) | ~ l3_lattices(k11_conlat_1(k7_conlat_2(X0))) | ~ v10_lattices(k11_conlat_1(k7_conlat_2(X0))) | v3_struct_0(k11_conlat_1(k7_conlat_2(X0))) | ~ l3_lattices(k1_lattice2(k11_conlat_1(X0))) | ~ v10_lattices(k1_lattice2(k11_conlat_1(X0))) | v3_struct_0(k1_lattice2(k11_conlat_1(X0)))), inference(resolution, [], [f927, f517])).
fof(f517, plain, ! [X0, X1] : (~ r1_filter_1(X0, X1) | r1_filter_1(X1, X0) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ! [X0, X1] : (r1_filter_1(X1, X0) | ~ r1_filter_1(X0, X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f205])).
fof(f205, plain, ! [X0, X1] : ((r1_filter_1(X1, X0) | ~ r1_filter_1(X0, X1)) | (~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f94])).
fof(f94, plain, ! [X0, X1] : ((l3_lattices(X1) & v10_lattices(X1) & ~ v3_struct_0(X1) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (r1_filter_1(X0, X1) => r1_filter_1(X1, X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', symmetry_r1_filter_1)).
fof(f927, plain, ! [X0] : (r1_filter_1(k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0))) | ~ l3_lattices(k11_conlat_1(k7_conlat_2(X0))) | ~ v10_lattices(k11_conlat_1(k7_conlat_2(X0))) | v3_struct_0(k11_conlat_1(k7_conlat_2(X0))) | ~ l3_lattices(k1_lattice2(k11_conlat_1(X0))) | ~ v10_lattices(k1_lattice2(k11_conlat_1(X0))) | v3_struct_0(k1_lattice2(k11_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(subsumption_resolution, [], [f924, f331])).
fof(f331, plain, ! [X0] : (m1_lattice4(k10_conlat_2(X0), k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ! [X0] : (m1_lattice4(k10_conlat_2(X0), k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f139])).
fof(f139, plain, ! [X0] : (m1_lattice4(k10_conlat_2(X0), k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => m1_lattice4(k10_conlat_2(X0), k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_k10_conlat_2)).
fof(f924, plain, ! [X0] : (r1_filter_1(k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0))) | ~ m1_lattice4(k10_conlat_2(X0), k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0))) | ~ l3_lattices(k11_conlat_1(k7_conlat_2(X0))) | ~ v10_lattices(k11_conlat_1(k7_conlat_2(X0))) | v3_struct_0(k11_conlat_1(k7_conlat_2(X0))) | ~ l3_lattices(k1_lattice2(k11_conlat_1(X0))) | ~ v10_lattices(k1_lattice2(k11_conlat_1(X0))) | v3_struct_0(k1_lattice2(k11_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(resolution, [], [f325, f519])).
fof(f519, plain, ! [X0] : (v3_lattice4(k10_conlat_2(X0), k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f209])).
fof(f209, plain, ! [X0] : (v3_lattice4(k10_conlat_2(X0), k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f208])).
fof(f208, plain, ! [X0] : (v3_lattice4(k10_conlat_2(X0), k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f96])).
fof(f96, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => v3_lattice4(k10_conlat_2(X0), k1_lattice2(k11_conlat_1(X0)), k11_conlat_1(k7_conlat_2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', t21_conlat_2)).
fof(f325, plain, ! [X2, X0, X1] : (~ v3_lattice4(X2, X0, X1) | r1_filter_1(X0, X1) | ~ m1_lattice4(X2, X0, X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f234])).
fof(f234, plain, ! [X0] : (! [X1] : (((r1_filter_1(X0, X1) | ! [X2] : (~ v3_lattice4(X2, X0, X1) | ~ m1_lattice4(X2, X0, X1))) & ((v3_lattice4(sK6(X0, X1), X0, X1) & m1_lattice4(sK6(X0, X1), X0, X1)) | ~ r1_filter_1(X0, X1))) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6])], [f232, f233])).
fof(f233, plain, ! [X1, X0] : (? [X3] : (v3_lattice4(X3, X0, X1) & m1_lattice4(X3, X0, X1)) => (v3_lattice4(sK6(X0, X1), X0, X1) & m1_lattice4(sK6(X0, X1), X0, X1))), introduced(choice_axiom, [])).
fof(f232, plain, ! [X0] : (! [X1] : (((r1_filter_1(X0, X1) | ! [X2] : (~ v3_lattice4(X2, X0, X1) | ~ m1_lattice4(X2, X0, X1))) & (? [X3] : (v3_lattice4(X3, X0, X1) & m1_lattice4(X3, X0, X1)) | ~ r1_filter_1(X0, X1))) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(rectify, [], [f231])).
fof(f231, plain, ! [X0] : (! [X1] : (((r1_filter_1(X0, X1) | ! [X2] : (~ v3_lattice4(X2, X0, X1) | ~ m1_lattice4(X2, X0, X1))) & (? [X2] : (v3_lattice4(X2, X0, X1) & m1_lattice4(X2, X0, X1)) | ~ r1_filter_1(X0, X1))) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f133])).
fof(f133, plain, ! [X0] : (! [X1] : ((r1_filter_1(X0, X1) <=> ? [X2] : (v3_lattice4(X2, X0, X1) & m1_lattice4(X2, X0, X1))) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f132])).
fof(f132, plain, ! [X0] : (! [X1] : ((r1_filter_1(X0, X1) <=> ? [X2] : (v3_lattice4(X2, X0, X1) & m1_lattice4(X2, X0, X1))) | (~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((l3_lattices(X1) & v10_lattices(X1) & ~ v3_struct_0(X1)) => (r1_filter_1(X0, X1) <=> ? [X2] : (v3_lattice4(X2, X0, X1) & m1_lattice4(X2, X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', d5_lattice4)).
fof(f3658, plain, spl33_222, inference(avatar_contradiction_clause, [], [f3657])).
fof(f3657, plain, ($false | spl33_222), inference(subsumption_resolution, [], [f3656, f294])).
fof(f3656, plain, (~ l2_conlat_1(sK5) | spl33_222), inference(subsumption_resolution, [], [f3655, f293])).
fof(f3655, plain, (v3_conlat_1(sK5) | ~ l2_conlat_1(sK5) | spl33_222), inference(resolution, [], [f3439, f607])).
fof(f607, plain, ! [X1] : (l1_lattices(k11_conlat_1(X1)) | v3_conlat_1(X1) | ~ l2_conlat_1(X1)), inference(resolution, [], [f334, f345])).
fof(f345, plain, ! [X0] : (~ l3_lattices(X0) | l1_lattices(X0)), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ! [X0] : ((l2_lattices(X0) & l1_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_l3_lattices)).
fof(f3439, plain, (~ l1_lattices(k11_conlat_1(sK5)) | spl33_222), inference(avatar_component_clause, [], [f3437])).
fof(f3437, plain, (spl33_222 <=> l1_lattices(k11_conlat_1(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl33_222])])).
fof(f3565, plain, spl33_221, inference(avatar_contradiction_clause, [], [f3564])).
fof(f3564, plain, ($false | spl33_221), inference(subsumption_resolution, [], [f3563, f294])).
fof(f3563, plain, (~ l2_conlat_1(sK5) | spl33_221), inference(subsumption_resolution, [], [f3562, f293])).
fof(f3562, plain, (v3_conlat_1(sK5) | ~ l2_conlat_1(sK5) | spl33_221), inference(resolution, [], [f3435, f606])).
fof(f606, plain, ! [X0] : (l2_lattices(k11_conlat_1(X0)) | v3_conlat_1(X0) | ~ l2_conlat_1(X0)), inference(resolution, [], [f334, f346])).
fof(f346, plain, ! [X0] : (~ l3_lattices(X0) | l2_lattices(X0)), inference(cnf_transformation, [], [f151])).
fof(f3435, plain, (~ l2_lattices(k11_conlat_1(sK5)) | spl33_221), inference(avatar_component_clause, [], [f3433])).
fof(f3433, plain, (spl33_221 <=> l2_lattices(k11_conlat_1(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl33_221])])).
fof(f3444, plain, (~ spl33_221 | ~ spl33_222 | spl33_223), inference(avatar_split_clause, [], [f3412, f3441, f3437, f3433])).
fof(f3412, plain, (l3_lattices(k11_conlat_1(sK5)) | ~ l1_lattices(k11_conlat_1(sK5)) | ~ l2_lattices(k11_conlat_1(sK5))), inference(superposition, [], [f1821, f1652])).
fof(f1652, plain, (k11_conlat_1(sK5) = g3_lattices(u1_struct_0(k11_conlat_1(sK5)), u2_lattices(k11_conlat_1(sK5)), u1_lattices(k11_conlat_1(sK5)))), inference(subsumption_resolution, [], [f1647, f293])).
fof(f1647, plain, ((k11_conlat_1(sK5) = g3_lattices(u1_struct_0(k11_conlat_1(sK5)), u2_lattices(k11_conlat_1(sK5)), u1_lattices(k11_conlat_1(sK5)))) | v3_conlat_1(sK5)), inference(resolution, [], [f814, f294])).
fof(f814, plain, ! [X0] : (~ l2_conlat_1(X0) | (k11_conlat_1(X0) = g3_lattices(u1_struct_0(k11_conlat_1(X0)), u2_lattices(k11_conlat_1(X0)), u1_lattices(k11_conlat_1(X0)))) | v3_conlat_1(X0)), inference(subsumption_resolution, [], [f803, f334])).
fof(f803, plain, ! [X0] : ((k11_conlat_1(X0) = g3_lattices(u1_struct_0(k11_conlat_1(X0)), u2_lattices(k11_conlat_1(X0)), u1_lattices(k11_conlat_1(X0)))) | ~ l3_lattices(k11_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(resolution, [], [f296, f333])).
fof(f333, plain, ! [X0] : (v3_lattices(k11_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f142])).
fof(f296, plain, ! [X0] : (~ v3_lattices(X0) | (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ! [X0] : ((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f112])).
fof(f112, plain, ! [X0] : (((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (l3_lattices(X0) => (v3_lattices(X0) => (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', abstractness_v3_lattices)).
fof(f1821, plain, ! [X0] : (l3_lattices(g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0))) | ~ l1_lattices(X0) | ~ l2_lattices(X0)), inference(subsumption_resolution, [], [f1820, f354])).
fof(f354, plain, ! [X0] : (v1_funct_1(u2_lattices(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ! [X0] : ((m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0))) | ~ l2_lattices(X0)), inference(ennf_transformation, [], [f41])).
fof(f41, plain, ! [X0] : (l2_lattices(X0) => (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_u2_lattices)).
fof(f1820, plain, ! [X0] : (l3_lattices(g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0))) | ~ v1_funct_1(u2_lattices(X0)) | ~ l1_lattices(X0) | ~ l2_lattices(X0)), inference(subsumption_resolution, [], [f1816, f801])).
fof(f801, plain, ! [X1] : (m1_relset_1(u2_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l2_lattices(X1)), inference(resolution, [], [f356, f513])).
fof(f513, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f292])).
fof(f292, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f91])).
fof(f91, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', redefinition_m2_relset_1)).
fof(f356, plain, ! [X0] : (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f156])).
fof(f1816, plain, ! [X0] : (~ m1_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | l3_lattices(g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0))) | ~ v1_funct_1(u2_lattices(X0)) | ~ l1_lattices(X0) | ~ l2_lattices(X0)), inference(resolution, [], [f945, f355])).
fof(f355, plain, ! [X0] : (v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f156])).
fof(f945, plain, ! [X2, X3] : (~ v1_funct_2(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | ~ m1_relset_1(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | l3_lattices(g3_lattices(u1_struct_0(X2), X3, u1_lattices(X2))) | ~ v1_funct_1(X3) | ~ l1_lattices(X2)), inference(subsumption_resolution, [], [f944, f351])).
fof(f351, plain, ! [X0] : (v1_funct_1(u1_lattices(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ! [X0] : ((m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0))) | ~ l1_lattices(X0)), inference(ennf_transformation, [], [f38])).
fof(f38, plain, ! [X0] : (l1_lattices(X0) => (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_u1_lattices)).
fof(f944, plain, ! [X2, X3] : (l3_lattices(g3_lattices(u1_struct_0(X2), X3, u1_lattices(X2))) | ~ v1_funct_1(u1_lattices(X2)) | ~ m1_relset_1(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | ~ v1_funct_2(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | ~ v1_funct_1(X3) | ~ l1_lattices(X2)), inference(subsumption_resolution, [], [f940, f793])).
fof(f793, plain, ! [X1] : (m1_relset_1(u1_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l1_lattices(X1)), inference(resolution, [], [f353, f513])).
fof(f353, plain, ! [X0] : (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f155])).
fof(f940, plain, ! [X2, X3] : (~ m1_relset_1(u1_lattices(X2), k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | l3_lattices(g3_lattices(u1_struct_0(X2), X3, u1_lattices(X2))) | ~ v1_funct_1(u1_lattices(X2)) | ~ m1_relset_1(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | ~ v1_funct_2(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | ~ v1_funct_1(X3) | ~ l1_lattices(X2)), inference(resolution, [], [f330, f352])).
fof(f352, plain, ! [X0] : (v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f155])).
fof(f330, plain, ! [X2, X0, X1] : (~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | l3_lattices(g3_lattices(X0, X1, X2)) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ! [X0, X1, X2] : ((l3_lattices(g3_lattices(X0, X1, X2)) & v3_lattices(g3_lattices(X0, X1, X2))) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(flattening, [], [f137])).
fof(f137, plain, ! [X0, X1, X2] : ((l3_lattices(g3_lattices(X0, X1, X2)) & v3_lattices(g3_lattices(X0, X1, X2))) | (~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X2) & m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X1)) => (l3_lattices(g3_lattices(X0, X1, X2)) & v3_lattices(g3_lattices(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_g3_lattices)).
fof(f3342, plain, ~ spl33_96, inference(avatar_contradiction_clause, [], [f3341])).
fof(f3341, plain, ($false | ~ spl33_96), inference(subsumption_resolution, [], [f3340, f293])).
fof(f3340, plain, (v3_conlat_1(sK5) | ~ spl33_96), inference(subsumption_resolution, [], [f3339, f294])).
fof(f3339, plain, (~ l2_conlat_1(sK5) | v3_conlat_1(sK5) | ~ spl33_96), inference(resolution, [], [f1615, f339])).
fof(f339, plain, ! [X0] : (~ v3_conlat_1(k7_conlat_2(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ! [X0] : ((l2_conlat_1(k7_conlat_2(X0)) & v4_conlat_1(k7_conlat_2(X0)) & ~ v3_conlat_1(k7_conlat_2(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f146])).
fof(f146, plain, ! [X0] : ((l2_conlat_1(k7_conlat_2(X0)) & v4_conlat_1(k7_conlat_2(X0)) & ~ v3_conlat_1(k7_conlat_2(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (l2_conlat_1(k7_conlat_2(X0)) & v4_conlat_1(k7_conlat_2(X0)) & ~ v3_conlat_1(k7_conlat_2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_k7_conlat_2)).
fof(f1615, plain, (v3_conlat_1(k7_conlat_2(sK5)) | ~ spl33_96), inference(avatar_component_clause, [], [f1613])).
fof(f3329, plain, spl33_90, inference(avatar_contradiction_clause, [], [f3328])).
fof(f3328, plain, ($false | spl33_90), inference(subsumption_resolution, [], [f3327, f294])).
fof(f3327, plain, (~ l2_conlat_1(sK5) | spl33_90), inference(resolution, [], [f3314, f711])).
fof(f711, plain, ! [X0] : (m1_relset_1(u3_conlat_1(X0), u1_conlat_1(X0), u2_conlat_1(X0)) | ~ l2_conlat_1(X0)), inference(resolution, [], [f357, f513])).
fof(f357, plain, ! [X0] : (m2_relset_1(u3_conlat_1(X0), u1_conlat_1(X0), u2_conlat_1(X0)) | ~ l2_conlat_1(X0)), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ! [X0] : (m2_relset_1(u3_conlat_1(X0), u1_conlat_1(X0), u2_conlat_1(X0)) | ~ l2_conlat_1(X0)), inference(ennf_transformation, [], [f42])).
fof(f42, plain, ! [X0] : (l2_conlat_1(X0) => m2_relset_1(u3_conlat_1(X0), u1_conlat_1(X0), u2_conlat_1(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_u3_conlat_1)).
fof(f3314, plain, (~ m1_relset_1(u3_conlat_1(sK5), u1_conlat_1(sK5), u2_conlat_1(sK5)) | spl33_90), inference(resolution, [], [f1589, f784])).
fof(f784, plain, ! [X4, X5, X3] : (m1_relset_1(k6_relset_1(X4, X5, X3), X5, X4) | ~ m1_relset_1(X3, X4, X5)), inference(resolution, [], [f338, f513])).
fof(f338, plain, ! [X2, X0, X1] : (m2_relset_1(k6_relset_1(X0, X1, X2), X1, X0) | ~ m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ! [X0, X1, X2] : (m2_relset_1(k6_relset_1(X0, X1, X2), X1, X0) | ~ m1_relset_1(X2, X0, X1)), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ! [X0, X1, X2] : (m1_relset_1(X2, X0, X1) => m2_relset_1(k6_relset_1(X0, X1, X2), X1, X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_k6_relset_1)).
fof(f1589, plain, (~ m1_relset_1(k6_relset_1(u1_conlat_1(sK5), u2_conlat_1(sK5), u3_conlat_1(sK5)), u2_conlat_1(sK5), u1_conlat_1(sK5)) | spl33_90), inference(avatar_component_clause, [], [f1587])).
fof(f1587, plain, (spl33_90 <=> m1_relset_1(k6_relset_1(u1_conlat_1(sK5), u2_conlat_1(sK5), u3_conlat_1(sK5)), u2_conlat_1(sK5), u1_conlat_1(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl33_90])])).
fof(f1630, plain, (spl33_96 | ~ spl33_90 | spl33_99), inference(avatar_split_clause, [], [f1584, f1627, f1587, f1613])).
fof(f1584, plain, (sP0(k7_conlat_2(sK5)) | ~ m1_relset_1(k6_relset_1(u1_conlat_1(sK5), u2_conlat_1(sK5), u3_conlat_1(sK5)), u2_conlat_1(sK5), u1_conlat_1(sK5)) | v3_conlat_1(k7_conlat_2(sK5))), inference(superposition, [], [f708, f896])).
fof(f896, plain, (k7_conlat_2(sK5) = g2_conlat_1(u2_conlat_1(sK5), u1_conlat_1(sK5), k6_relset_1(u1_conlat_1(sK5), u2_conlat_1(sK5), u3_conlat_1(sK5)))), inference(subsumption_resolution, [], [f891, f293])).
fof(f891, plain, ((k7_conlat_2(sK5) = g2_conlat_1(u2_conlat_1(sK5), u1_conlat_1(sK5), k6_relset_1(u1_conlat_1(sK5), u2_conlat_1(sK5), u3_conlat_1(sK5)))) | v3_conlat_1(sK5)), inference(resolution, [], [f326, f294])).
fof(f326, plain, ! [X0] : (~ l2_conlat_1(X0) | (k7_conlat_2(X0) = g2_conlat_1(u2_conlat_1(X0), u1_conlat_1(X0), k6_relset_1(u1_conlat_1(X0), u2_conlat_1(X0), u3_conlat_1(X0)))) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ! [X0] : ((k7_conlat_2(X0) = g2_conlat_1(u2_conlat_1(X0), u1_conlat_1(X0), k6_relset_1(u1_conlat_1(X0), u2_conlat_1(X0), u3_conlat_1(X0)))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f134])).
fof(f134, plain, ! [X0] : ((k7_conlat_2(X0) = g2_conlat_1(u2_conlat_1(X0), u1_conlat_1(X0), k6_relset_1(u1_conlat_1(X0), u2_conlat_1(X0), u3_conlat_1(X0)))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (k7_conlat_2(X0) = g2_conlat_1(u2_conlat_1(X0), u1_conlat_1(X0), k6_relset_1(u1_conlat_1(X0), u2_conlat_1(X0), u3_conlat_1(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', d7_conlat_2)).
fof(f708, plain, ! [X6, X8, X7] : (sP0(g2_conlat_1(X7, X8, X6)) | ~ m1_relset_1(X6, X7, X8) | v3_conlat_1(g2_conlat_1(X7, X8, X6))), inference(resolution, [], [f328, f382])).
fof(f382, plain, ! [X0] : (~ l2_conlat_1(X0) | sP0(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f220])).
fof(f220, plain, ! [X0] : (sP0(X0) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(definition_folding, [], [f163, e219])).
fof(f163, plain, ! [X0] : ((v4_lattice3(k11_conlat_1(X0)) & v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f162])).
fof(f162, plain, ! [X0] : ((v4_lattice3(k11_conlat_1(X0)) & v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (v4_lattice3(k11_conlat_1(X0)) & v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', fc1_conlat_2)).
fof(f328, plain, ! [X2, X0, X1] : (l2_conlat_1(g2_conlat_1(X0, X1, X2)) | ~ m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ! [X0, X1, X2] : ((l2_conlat_1(g2_conlat_1(X0, X1, X2)) & v4_conlat_1(g2_conlat_1(X0, X1, X2))) | ~ m1_relset_1(X2, X0, X1)), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0, X1, X2] : (m1_relset_1(X2, X0, X1) => (l2_conlat_1(g2_conlat_1(X0, X1, X2)) & v4_conlat_1(g2_conlat_1(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT346+1.p', dt_g2_conlat_1)).
fof(f1599, plain, (~ spl33_90 | spl33_92), inference(avatar_split_clause, [], [f1575, f1596, f1587])).
fof(f1575, plain, (l2_conlat_1(k7_conlat_2(sK5)) | ~ m1_relset_1(k6_relset_1(u1_conlat_1(sK5), u2_conlat_1(sK5), u3_conlat_1(sK5)), u2_conlat_1(sK5), u1_conlat_1(sK5))), inference(superposition, [], [f328, f896])).