fof(f32091, plain, $false, inference(avatar_sat_refutation, [], [f1129, f1130, f1411, f2846, f3029, f3253, f3268, f3271, f3274, f3319, f3328, f3362, f4210, f4763, f7113, f7443, f16555, f16670, f16860, f16928, f16943, f16961, f17071, f32053, f32072, f32090])).
fof(f32090, plain, (~ spl50_1 | spl50_7 | ~ spl50_72 | spl50_1819), inference(avatar_contradiction_clause, [], [f32089])).
fof(f32089, plain, ($false | (~ spl50_1 | spl50_7 | ~ spl50_72 | spl50_1819)), inference(subsumption_resolution, [], [f32088, f3219])).
fof(f3219, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl50_7), inference(subsumption_resolution, [], [f3218, f1329])).
fof(f1329, plain, (~ v3_struct_0(sK9) | spl50_7), inference(avatar_component_clause, [], [f1328])).
fof(f1328, plain, (spl50_7 <=> v3_struct_0(sK9)), introduced(avatar_definition, [new_symbols(naming, [spl50_7])])).
fof(f3218, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3217, f571])).
fof(f571, plain, l1_orders_2(sK9), inference(cnf_transformation, [], [f480])).
fof(f480, plain, (((((((~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK14) | ~ r1_waybel_3(sK9, sK11, sK12)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK14) | r1_waybel_3(sK9, sK11, sK12)) & (sK12 = sK14) & (sK11 = sK13) & m1_subset_1(sK14, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(sK13, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(sK12, u1_struct_0(sK9))) & m1_subset_1(sK11, u1_struct_0(sK9))) & v22_waybel_0(sK10, sK9, sK9) & m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v8_waybel_1(sK10, sK9) & v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(sK10)) & l1_orders_2(sK9) & v3_lattice3(sK9) & v2_lattice3(sK9) & v1_lattice3(sK9) & v4_orders_2(sK9) & v3_orders_2(sK9) & v2_orders_2(sK9)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10, sK11, sK12, sK13, sK14])], [f473, f479, f478, f477, f476, f475, f474])).
fof(f474, plain, (? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : ((~ r1_waybel_3(k2_yellow_2(X0, X0, X1), X4, X5) | ~ r1_waybel_3(X0, X2, X3)) & (r1_waybel_3(k2_yellow_2(X0, X0, X1), X4, X5) | r1_waybel_3(X0, X2, X3)) & (X3 = X5) & (X2 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(X0, X0, X1)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(X0, X0, X1)))) & m1_subset_1(X3, u1_struct_0(X0))) & m1_subset_1(X2, u1_struct_0(X0))) & v22_waybel_0(X1, X0, X0) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, X1), X4, X5) | ~ r1_waybel_3(sK9, X2, X3)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, X1), X4, X5) | r1_waybel_3(sK9, X2, X3)) & (X3 = X5) & (X2 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(sK9, sK9, X1)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(sK9, sK9, X1)))) & m1_subset_1(X3, u1_struct_0(sK9))) & m1_subset_1(X2, u1_struct_0(sK9))) & v22_waybel_0(X1, sK9, sK9) & m2_relset_1(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v8_waybel_1(X1, sK9) & v1_funct_2(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(X1)) & l1_orders_2(sK9) & v3_lattice3(sK9) & v2_lattice3(sK9) & v1_lattice3(sK9) & v4_orders_2(sK9) & v3_orders_2(sK9) & v2_orders_2(sK9))), introduced(choice_axiom, [])).
fof(f475, plain, (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, X1), X4, X5) | ~ r1_waybel_3(sK9, X2, X3)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, X1), X4, X5) | r1_waybel_3(sK9, X2, X3)) & (X3 = X5) & (X2 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(sK9, sK9, X1)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(sK9, sK9, X1)))) & m1_subset_1(X3, u1_struct_0(sK9))) & m1_subset_1(X2, u1_struct_0(sK9))) & v22_waybel_0(X1, sK9, sK9) & m2_relset_1(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v8_waybel_1(X1, sK9) & v1_funct_2(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(X1)) => (? [X2] : (? [X3] : (? [X4] : (? [X5] : ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | ~ r1_waybel_3(sK9, X2, X3)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | r1_waybel_3(sK9, X2, X3)) & (X3 = X5) & (X2 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(X3, u1_struct_0(sK9))) & m1_subset_1(X2, u1_struct_0(sK9))) & v22_waybel_0(sK10, sK9, sK9) & m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v8_waybel_1(sK10, sK9) & v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(sK10))), introduced(choice_axiom, [])).
fof(f476, plain, (? [X2] : (? [X3] : (? [X4] : (? [X5] : ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | ~ r1_waybel_3(sK9, X2, X3)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | r1_waybel_3(sK9, X2, X3)) & (X3 = X5) & (X2 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(X3, u1_struct_0(sK9))) & m1_subset_1(X2, u1_struct_0(sK9))) => (? [X3] : (? [X4] : (? [X5] : ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | ~ r1_waybel_3(sK9, sK11, X3)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | r1_waybel_3(sK9, sK11, X3)) & (X3 = X5) & (sK11 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(X3, u1_struct_0(sK9))) & m1_subset_1(sK11, u1_struct_0(sK9)))), introduced(choice_axiom, [])).
fof(f477, plain, (? [X3] : (? [X4] : (? [X5] : ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | ~ r1_waybel_3(sK9, sK11, X3)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | r1_waybel_3(sK9, sK11, X3)) & (X3 = X5) & (sK11 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(X3, u1_struct_0(sK9))) => (? [X4] : (? [X5] : ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | ~ r1_waybel_3(sK9, sK11, sK12)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | r1_waybel_3(sK9, sK11, sK12)) & (sK12 = X5) & (sK11 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(sK12, u1_struct_0(sK9)))), introduced(choice_axiom, [])).
fof(f478, plain, (? [X4] : (? [X5] : ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | ~ r1_waybel_3(sK9, sK11, sK12)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X4, X5) | r1_waybel_3(sK9, sK11, sK12)) & (sK12 = X5) & (sK11 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) => (? [X5] : ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, X5) | ~ r1_waybel_3(sK9, sK11, sK12)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, X5) | r1_waybel_3(sK9, sK11, sK12)) & (sK12 = X5) & (sK11 = sK13) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) & m1_subset_1(sK13, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))))), introduced(choice_axiom, [])).
fof(f479, plain, (? [X5] : ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, X5) | ~ r1_waybel_3(sK9, sK11, sK12)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, X5) | r1_waybel_3(sK9, sK11, sK12)) & (sK12 = X5) & (sK11 = sK13) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))) => ((~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK14) | ~ r1_waybel_3(sK9, sK11, sK12)) & (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK14) | r1_waybel_3(sK9, sK11, sK12)) & (sK12 = sK14) & (sK11 = sK13) & m1_subset_1(sK14, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))))), introduced(choice_axiom, [])).
fof(f473, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : ((~ r1_waybel_3(k2_yellow_2(X0, X0, X1), X4, X5) | ~ r1_waybel_3(X0, X2, X3)) & (r1_waybel_3(k2_yellow_2(X0, X0, X1), X4, X5) | r1_waybel_3(X0, X2, X3)) & (X3 = X5) & (X2 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(X0, X0, X1)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(X0, X0, X1)))) & m1_subset_1(X3, u1_struct_0(X0))) & m1_subset_1(X2, u1_struct_0(X0))) & v22_waybel_0(X1, X0, X0) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(flattening, [], [f472])).
fof(f472, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (((~ r1_waybel_3(k2_yellow_2(X0, X0, X1), X4, X5) | ~ r1_waybel_3(X0, X2, X3)) & (r1_waybel_3(k2_yellow_2(X0, X0, X1), X4, X5) | r1_waybel_3(X0, X2, X3))) & (X3 = X5) & (X2 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(X0, X0, X1)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(X0, X0, X1)))) & m1_subset_1(X3, u1_struct_0(X0))) & m1_subset_1(X2, u1_struct_0(X0))) & v22_waybel_0(X1, X0, X0) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(nnf_transformation, [], [f238])).
fof(f238, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (~ (r1_waybel_3(X0, X2, X3) <=> r1_waybel_3(k2_yellow_2(X0, X0, X1), X4, X5)) & (X3 = X5) & (X2 = X4) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(X0, X0, X1)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(X0, X0, X1)))) & m1_subset_1(X3, u1_struct_0(X0))) & m1_subset_1(X2, u1_struct_0(X0))) & v22_waybel_0(X1, X0, X0) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(flattening, [], [f237])).
fof(f237, plain, ? [X0] : (? [X1] : ((? [X2] : (? [X3] : (? [X4] : (? [X5] : ((~ (r1_waybel_3(X0, X2, X3) <=> r1_waybel_3(k2_yellow_2(X0, X0, X1), X4, X5)) & ((X3 = X5) & (X2 = X4))) & m1_subset_1(X5, u1_struct_0(k2_yellow_2(X0, X0, X1)))) & m1_subset_1(X4, u1_struct_0(k2_yellow_2(X0, X0, X1)))) & m1_subset_1(X3, u1_struct_0(X0))) & m1_subset_1(X2, u1_struct_0(X0))) & v22_waybel_0(X1, X0, X0)) & (m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1))) & (l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v22_waybel_0(X1, X0, X0) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3, u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4, u1_struct_0(k2_yellow_2(X0, X0, X1))) => ! [X5] : (m1_subset_1(X5, u1_struct_0(k2_yellow_2(X0, X0, X1))) => (((X3 = X5) & (X2 = X4)) => (r1_waybel_3(X0, X2, X3) <=> r1_waybel_3(k2_yellow_2(X0, X0, X1), X4, X5)))))))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v22_waybel_0(X1, X0, X0) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3, u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4, u1_struct_0(k2_yellow_2(X0, X0, X1))) => ! [X5] : (m1_subset_1(X5, u1_struct_0(k2_yellow_2(X0, X0, X1))) => (((X3 = X5) & (X2 = X4)) => (r1_waybel_3(X0, X2, X3) <=> r1_waybel_3(k2_yellow_2(X0, X0, X1), X4, X5)))))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', t37_waybel34)).
fof(f3217, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3216, f572])).
fof(f572, plain, v1_funct_1(sK10), inference(cnf_transformation, [], [f480])).
fof(f3216, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3215, f573])).
fof(f573, plain, v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)), inference(cnf_transformation, [], [f480])).
fof(f3215, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f3196])).
fof(f3196, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f750, f1467])).
fof(f1467, plain, m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)), inference(resolution, [], [f1097, f575])).
fof(f575, plain, m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)), inference(cnf_transformation, [], [f480])).
fof(f1097, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f563])).
fof(f563, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f145])).
fof(f145, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', redefinition_m2_relset_1)).
fof(f750, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f336])).
fof(f336, plain, ! [X0, X1, X2] : ((m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f335])).
fof(f335, plain, ! [X0, X1, X2] : ((m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f69])).
fof(f69, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', dt_k2_yellow_2)).
fof(f32088, plain, (~ v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | (~ spl50_1 | spl50_7 | ~ spl50_72 | spl50_1819)), inference(subsumption_resolution, [], [f32087, f3968])).
fof(f3968, plain, (v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl50_7), inference(subsumption_resolution, [], [f3967, f1329])).
fof(f3967, plain, (v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3966, f565])).
fof(f565, plain, v2_orders_2(sK9), inference(cnf_transformation, [], [f480])).
fof(f3966, plain, (v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3965, f566])).
fof(f566, plain, v3_orders_2(sK9), inference(cnf_transformation, [], [f480])).
fof(f3965, plain, (v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3964, f567])).
fof(f567, plain, v4_orders_2(sK9), inference(cnf_transformation, [], [f480])).
fof(f3964, plain, (v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3963, f571])).
fof(f3963, plain, (v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3962, f572])).
fof(f3962, plain, (v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3961, f573])).
fof(f3961, plain, (v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3954, f574])).
fof(f574, plain, v8_waybel_1(sK10, sK9), inference(cnf_transformation, [], [f480])).
fof(f3954, plain, (~ v8_waybel_1(sK10, sK9) | v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f1114, f575])).
fof(f1114, plain, ! [X0, X1] : (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v8_waybel_1(X1, X0) | v8_yellow_0(k2_yellow_2(X0, X0, X1), X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f449])).
fof(f449, plain, ! [X0] : (! [X1] : (((v8_yellow_0(k2_yellow_2(X0, X0, X1), X0) | ~ v8_waybel_1(X1, X0)) & (v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) | ~ v7_waybel_1(X1, X0))) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f448])).
fof(f448, plain, ! [X0] : (! [X1] : (((v8_yellow_0(k2_yellow_2(X0, X0, X1), X0) | ~ v8_waybel_1(X1, X0)) & (v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) | ~ v7_waybel_1(X1, X0))) | (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | (~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f156])).
fof(f156, plain, ! [X0] : ((l1_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => ((v8_waybel_1(X1, X0) => v8_yellow_0(k2_yellow_2(X0, X0, X1), X0)) & (v7_waybel_1(X1, X0) => v7_yellow_0(k2_yellow_2(X0, X0, X1), X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', t56_waybel_1)).
fof(f32087, plain, (~ v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | (~ spl50_1 | ~ spl50_72 | spl50_1819)), inference(subsumption_resolution, [], [f32086, f3247])).
fof(f3247, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ spl50_72), inference(avatar_component_clause, [], [f3246])).
fof(f3246, plain, (spl50_72 <=> m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9)), introduced(avatar_definition, [new_symbols(naming, [spl50_72])])).
fof(f32086, plain, (~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | (~ spl50_1 | spl50_1819)), inference(subsumption_resolution, [], [f32085, f1132])).
fof(f1132, plain, m1_subset_1(sK11, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))), inference(forward_demodulation, [], [f579, f581])).
fof(f581, plain, (sK11 = sK13), inference(cnf_transformation, [], [f480])).
fof(f579, plain, m1_subset_1(sK13, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))), inference(cnf_transformation, [], [f480])).
fof(f32085, plain, (~ m1_subset_1(sK11, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | (~ spl50_1 | spl50_1819)), inference(subsumption_resolution, [], [f32083, f32068])).
fof(f32068, plain, (~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK11, sK12) | spl50_1819), inference(avatar_component_clause, [], [f32066])).
fof(f32066, plain, (spl50_1819 <=> r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK11, sK12)), introduced(avatar_definition, [new_symbols(naming, [spl50_1819])])).
fof(f32083, plain, (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK11, sK12) | ~ m1_subset_1(sK11, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v8_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ spl50_1), inference(resolution, [], [f32082, f1131])).
fof(f1131, plain, m1_subset_1(sK12, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))), inference(forward_demodulation, [], [f580, f582])).
fof(f582, plain, (sK12 = sK14), inference(cnf_transformation, [], [f480])).
fof(f580, plain, m1_subset_1(sK14, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))), inference(cnf_transformation, [], [f480])).
fof(f32082, plain, (! [X0] : (~ m1_subset_1(sK12, u1_struct_0(X0)) | r1_waybel_3(X0, sK11, sK12) | ~ m1_subset_1(sK11, u1_struct_0(X0)) | ~ m1_yellow_0(X0, sK9) | ~ v8_yellow_0(X0, sK9) | ~ v4_yellow_0(X0, sK9)) | ~ spl50_1), inference(subsumption_resolution, [], [f32081, f565])).
fof(f32081, plain, (! [X0] : (r1_waybel_3(X0, sK11, sK12) | ~ m1_subset_1(sK12, u1_struct_0(X0)) | ~ m1_subset_1(sK11, u1_struct_0(X0)) | ~ m1_yellow_0(X0, sK9) | ~ v8_yellow_0(X0, sK9) | ~ v4_yellow_0(X0, sK9) | ~ v2_orders_2(sK9)) | ~ spl50_1), inference(subsumption_resolution, [], [f32080, f566])).
fof(f32080, plain, (! [X0] : (r1_waybel_3(X0, sK11, sK12) | ~ m1_subset_1(sK12, u1_struct_0(X0)) | ~ m1_subset_1(sK11, u1_struct_0(X0)) | ~ m1_yellow_0(X0, sK9) | ~ v8_yellow_0(X0, sK9) | ~ v4_yellow_0(X0, sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)) | ~ spl50_1), inference(subsumption_resolution, [], [f32079, f567])).
fof(f32079, plain, (! [X0] : (r1_waybel_3(X0, sK11, sK12) | ~ m1_subset_1(sK12, u1_struct_0(X0)) | ~ m1_subset_1(sK11, u1_struct_0(X0)) | ~ m1_yellow_0(X0, sK9) | ~ v8_yellow_0(X0, sK9) | ~ v4_yellow_0(X0, sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)) | ~ spl50_1), inference(subsumption_resolution, [], [f32078, f568])).
fof(f568, plain, v1_lattice3(sK9), inference(cnf_transformation, [], [f480])).
fof(f32078, plain, (! [X0] : (r1_waybel_3(X0, sK11, sK12) | ~ m1_subset_1(sK12, u1_struct_0(X0)) | ~ m1_subset_1(sK11, u1_struct_0(X0)) | ~ m1_yellow_0(X0, sK9) | ~ v8_yellow_0(X0, sK9) | ~ v4_yellow_0(X0, sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)) | ~ spl50_1), inference(subsumption_resolution, [], [f32077, f569])).
fof(f569, plain, v2_lattice3(sK9), inference(cnf_transformation, [], [f480])).
fof(f32077, plain, (! [X0] : (r1_waybel_3(X0, sK11, sK12) | ~ m1_subset_1(sK12, u1_struct_0(X0)) | ~ m1_subset_1(sK11, u1_struct_0(X0)) | ~ m1_yellow_0(X0, sK9) | ~ v8_yellow_0(X0, sK9) | ~ v4_yellow_0(X0, sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)) | ~ spl50_1), inference(subsumption_resolution, [], [f32076, f570])).
fof(f570, plain, v3_lattice3(sK9), inference(cnf_transformation, [], [f480])).
fof(f32076, plain, (! [X0] : (r1_waybel_3(X0, sK11, sK12) | ~ m1_subset_1(sK12, u1_struct_0(X0)) | ~ m1_subset_1(sK11, u1_struct_0(X0)) | ~ m1_yellow_0(X0, sK9) | ~ v8_yellow_0(X0, sK9) | ~ v4_yellow_0(X0, sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)) | ~ spl50_1), inference(subsumption_resolution, [], [f32075, f571])).
fof(f32075, plain, (! [X0] : (r1_waybel_3(X0, sK11, sK12) | ~ m1_subset_1(sK12, u1_struct_0(X0)) | ~ m1_subset_1(sK11, u1_struct_0(X0)) | ~ m1_yellow_0(X0, sK9) | ~ v8_yellow_0(X0, sK9) | ~ v4_yellow_0(X0, sK9) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)) | ~ spl50_1), inference(subsumption_resolution, [], [f32074, f577])).
fof(f577, plain, m1_subset_1(sK11, u1_struct_0(sK9)), inference(cnf_transformation, [], [f480])).
fof(f32074, plain, (! [X0] : (r1_waybel_3(X0, sK11, sK12) | ~ m1_subset_1(sK12, u1_struct_0(X0)) | ~ m1_subset_1(sK11, u1_struct_0(X0)) | ~ m1_subset_1(sK11, u1_struct_0(sK9)) | ~ m1_yellow_0(X0, sK9) | ~ v8_yellow_0(X0, sK9) | ~ v4_yellow_0(X0, sK9) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)) | ~ spl50_1), inference(subsumption_resolution, [], [f32073, f578])).
fof(f578, plain, m1_subset_1(sK12, u1_struct_0(sK9)), inference(cnf_transformation, [], [f480])).
fof(f32073, plain, (! [X0] : (r1_waybel_3(X0, sK11, sK12) | ~ m1_subset_1(sK12, u1_struct_0(X0)) | ~ m1_subset_1(sK11, u1_struct_0(X0)) | ~ m1_subset_1(sK12, u1_struct_0(sK9)) | ~ m1_subset_1(sK11, u1_struct_0(sK9)) | ~ m1_yellow_0(X0, sK9) | ~ v8_yellow_0(X0, sK9) | ~ v4_yellow_0(X0, sK9) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)) | ~ spl50_1), inference(resolution, [], [f1123, f1139])).
fof(f1139, plain, ! [X4, X0, X5, X1] : (~ r1_waybel_3(X0, X4, X5) | r1_waybel_3(X1, X4, X5) | ~ m1_subset_1(X5, u1_struct_0(X1)) | ~ m1_subset_1(X4, u1_struct_0(X1)) | ~ m1_subset_1(X5, u1_struct_0(X0)) | ~ m1_subset_1(X4, u1_struct_0(X0)) | ~ m1_yellow_0(X1, X0) | ~ v8_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(subsumption_resolution, [], [f1120, f696])).
fof(f696, plain, ! [X0, X1] : (~ v8_yellow_0(X1, X0) | ~ v3_struct_0(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f294])).
fof(f294, plain, ! [X0] : (! [X1] : (~ v3_struct_0(X1) | ~ v8_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f293])).
fof(f293, plain, ! [X0] : (! [X1] : ((~ v3_struct_0(X1) | ~ v8_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f44])).
fof(f44, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v8_yellow_0(X1, X0) => ~ v3_struct_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', cc3_waybel10)).
fof(f1120, plain, ! [X4, X0, X5, X1] : (r1_waybel_3(X1, X4, X5) | ~ r1_waybel_3(X0, X4, X5) | ~ m1_subset_1(X5, u1_struct_0(X1)) | ~ m1_subset_1(X4, u1_struct_0(X1)) | ~ m1_subset_1(X5, u1_struct_0(X0)) | ~ m1_subset_1(X4, u1_struct_0(X0)) | ~ m1_yellow_0(X1, X0) | ~ v8_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(equality_resolution, [], [f1119])).
fof(f1119, plain, ! [X4, X2, X0, X5, X1] : (r1_waybel_3(X1, X4, X5) | ~ r1_waybel_3(X0, X2, X5) | ~ (X2 = X4) | ~ m1_subset_1(X5, u1_struct_0(X1)) | ~ m1_subset_1(X4, u1_struct_0(X1)) | ~ m1_subset_1(X5, u1_struct_0(X0)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_yellow_0(X1, X0) | ~ v8_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(equality_resolution, [], [f1110])).
fof(f1110, plain, ! [X4, X2, X0, X5, X3, X1] : (r1_waybel_3(X1, X4, X5) | ~ r1_waybel_3(X0, X2, X3) | ~ (X3 = X5) | ~ (X2 = X4) | ~ m1_subset_1(X5, u1_struct_0(X1)) | ~ m1_subset_1(X4, u1_struct_0(X1)) | ~ m1_subset_1(X3, u1_struct_0(X0)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_yellow_0(X1, X0) | ~ v8_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f444])).
fof(f444, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (r1_waybel_3(X1, X4, X5) | ~ r1_waybel_3(X0, X2, X3) | ~ (X3 = X5) | ~ (X2 = X4) | ~ m1_subset_1(X5, u1_struct_0(X1))) | ~ m1_subset_1(X4, u1_struct_0(X1))) | ~ m1_subset_1(X3, u1_struct_0(X0))) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ m1_yellow_0(X1, X0) | ~ v8_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f443])).
fof(f443, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : ((r1_waybel_3(X1, X4, X5) | (~ r1_waybel_3(X0, X2, X3) | ~ (X3 = X5) | ~ (X2 = X4))) | ~ m1_subset_1(X5, u1_struct_0(X1))) | ~ m1_subset_1(X4, u1_struct_0(X1))) | ~ m1_subset_1(X3, u1_struct_0(X0))) | ~ m1_subset_1(X2, u1_struct_0(X0))) | (~ m1_yellow_0(X1, X0) | ~ v8_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f153])).
fof(f153, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m1_yellow_0(X1, X0) & v8_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & ~ v3_struct_0(X1)) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3, u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4, u1_struct_0(X1)) => ! [X5] : (m1_subset_1(X5, u1_struct_0(X1)) => ((r1_waybel_3(X0, X2, X3) & (X3 = X5) & (X2 = X4)) => r1_waybel_3(X1, X4, X5)))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', t36_waybel34)).
fof(f1123, plain, (r1_waybel_3(sK9, sK11, sK12) | ~ spl50_1), inference(avatar_component_clause, [], [f1122])).
fof(f1122, plain, (spl50_1 <=> r1_waybel_3(sK9, sK11, sK12)), introduced(avatar_definition, [new_symbols(naming, [spl50_1])])).
fof(f32072, plain, (~ spl50_1819 | spl50_2), inference(avatar_split_clause, [], [f32071, f1126, f32066])).
fof(f1126, plain, (spl50_2 <=> r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK14)), introduced(avatar_definition, [new_symbols(naming, [spl50_2])])).
fof(f32071, plain, (~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK11, sK12) | spl50_2), inference(forward_demodulation, [], [f32070, f581])).
fof(f32070, plain, (~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK12) | spl50_2), inference(forward_demodulation, [], [f1128, f582])).
fof(f1128, plain, (~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK14) | spl50_2), inference(avatar_component_clause, [], [f1126])).
fof(f32053, plain, (spl50_1 | ~ spl50_2 | spl50_7 | ~ spl50_112), inference(avatar_contradiction_clause, [], [f32052])).
fof(f32052, plain, ($false | (spl50_1 | ~ spl50_2 | spl50_7 | ~ spl50_112)), inference(subsumption_resolution, [], [f32051, f1141])).
fof(f1141, plain, (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK11, sK12) | ~ spl50_2), inference(forward_demodulation, [], [f1140, f581])).
fof(f1140, plain, (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK12) | ~ spl50_2), inference(forward_demodulation, [], [f1127, f582])).
fof(f1127, plain, (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK14) | ~ spl50_2), inference(avatar_component_clause, [], [f1126])).
fof(f32051, plain, (~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK11, sK12) | (spl50_1 | spl50_7 | ~ spl50_112)), inference(subsumption_resolution, [], [f32050, f1132])).
fof(f32050, plain, (~ m1_subset_1(sK11, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK11, sK12) | (spl50_1 | spl50_7 | ~ spl50_112)), inference(subsumption_resolution, [], [f32038, f1124])).
fof(f1124, plain, (~ r1_waybel_3(sK9, sK11, sK12) | spl50_1), inference(avatar_component_clause, [], [f1122])).
fof(f32038, plain, (r1_waybel_3(sK9, sK11, sK12) | ~ m1_subset_1(sK11, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK11, sK12) | (spl50_7 | ~ spl50_112)), inference(superposition, [], [f7112, f4749])).
fof(f4749, plain, ((sK11 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK11)) | spl50_7), inference(subsumption_resolution, [], [f4748, f1329])).
fof(f4748, plain, ((sK11 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK11)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f4747, f571])).
fof(f4747, plain, ((sK11 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK11)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f4746, f572])).
fof(f4746, plain, ((sK11 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK11)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f4745, f573])).
fof(f4745, plain, ((sK11 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK11)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f4738, f575])).
fof(f4738, plain, ((sK11 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK11)) | ~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f4733])).
fof(f4733, plain, ((sK11 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK11)) | ~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f1109, f1132])).
fof(f1109, plain, ! [X2, X0, X3, X1] : (~ m1_subset_1(X3, u1_struct_0(k2_yellow_2(X0, X1, X2))) | (k7_yellow_2(u1_struct_0(k2_yellow_2(X0, X1, X2)), X1, k3_waybel_1(X0, X1, X2), X3) = X3) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f442])).
fof(f442, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((k7_yellow_2(u1_struct_0(k2_yellow_2(X0, X1, X2)), X1, k3_waybel_1(X0, X1, X2), X3) = X3) | ~ m1_subset_1(X3, u1_struct_0(k2_yellow_2(X0, X1, X2)))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | v3_struct_0(X1)) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f441])).
fof(f441, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((k7_yellow_2(u1_struct_0(k2_yellow_2(X0, X1, X2)), X1, k3_waybel_1(X0, X1, X2), X3) = X3) | ~ m1_subset_1(X3, u1_struct_0(k2_yellow_2(X0, X1, X2)))) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | v3_struct_0(X1))) | (~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f152])).
fof(f152, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((l1_orders_2(X1) & ~ v3_struct_0(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => ! [X3] : (m1_subset_1(X3, u1_struct_0(k2_yellow_2(X0, X1, X2))) => (k7_yellow_2(u1_struct_0(k2_yellow_2(X0, X1, X2)), X1, k3_waybel_1(X0, X1, X2), X3) = X3))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', t34_waybel_1)).
fof(f7112, plain, (! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12)) | ~ spl50_112), inference(avatar_component_clause, [], [f7111])).
fof(f7111, plain, (spl50_112 <=> ! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12))), introduced(avatar_definition, [new_symbols(naming, [spl50_112])])).
fof(f17071, plain, spl50_124, inference(avatar_contradiction_clause, [], [f17070])).
fof(f17070, plain, ($false | spl50_124), inference(subsumption_resolution, [], [f17069, f565])).
fof(f17069, plain, (~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17068, f566])).
fof(f17068, plain, (~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17067, f567])).
fof(f17067, plain, (~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17066, f568])).
fof(f17066, plain, (~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17065, f569])).
fof(f17065, plain, (~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17064, f570])).
fof(f17064, plain, (~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17063, f571])).
fof(f17063, plain, (~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17062, f572])).
fof(f17062, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17061, f573])).
fof(f17061, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17060, f574])).
fof(f17060, plain, (~ v8_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17059, f575])).
fof(f17059, plain, (~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v8_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(subsumption_resolution, [], [f17058, f576])).
fof(f576, plain, v22_waybel_0(sK10, sK9, sK9), inference(cnf_transformation, [], [f480])).
fof(f17058, plain, (~ v22_waybel_0(sK10, sK9, sK9) | ~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v8_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_124), inference(resolution, [], [f7442, f1107])).
fof(f1107, plain, ! [X0, X1] : (v22_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) | ~ v22_waybel_0(X1, X0, X0) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f564])).
fof(f564, plain, ! [X0] : (! [X1] : (((v22_waybel_0(X1, X0, X0) | ~ v22_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) & (v22_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) | ~ v22_waybel_0(X1, X0, X0))) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(nnf_transformation, [], [f440])).
fof(f440, plain, ! [X0] : (! [X1] : ((v22_waybel_0(X1, X0, X0) <=> v22_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f439])).
fof(f439, plain, ! [X0] : (! [X1] : ((v22_waybel_0(X1, X0, X0) <=> v22_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) | (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f151])).
fof(f151, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v22_waybel_0(X1, X0, X0) <=> v22_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', t34_waybel34)).
fof(f7442, plain, (~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | spl50_124), inference(avatar_component_clause, [], [f7440])).
fof(f7440, plain, (spl50_124 <=> v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl50_124])])).
fof(f16961, plain, spl50_118, inference(avatar_contradiction_clause, [], [f16960])).
fof(f16960, plain, ($false | spl50_118), inference(subsumption_resolution, [], [f16959, f565])).
fof(f16959, plain, (~ v2_orders_2(sK9) | spl50_118), inference(subsumption_resolution, [], [f16958, f566])).
fof(f16958, plain, (~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_118), inference(subsumption_resolution, [], [f16957, f567])).
fof(f16957, plain, (~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_118), inference(subsumption_resolution, [], [f16956, f568])).
fof(f16956, plain, (~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_118), inference(subsumption_resolution, [], [f16955, f569])).
fof(f16955, plain, (~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_118), inference(subsumption_resolution, [], [f16954, f570])).
fof(f16954, plain, (~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_118), inference(subsumption_resolution, [], [f16953, f571])).
fof(f16953, plain, (~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_118), inference(subsumption_resolution, [], [f16952, f572])).
fof(f16952, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_118), inference(subsumption_resolution, [], [f16951, f573])).
fof(f16951, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_118), inference(subsumption_resolution, [], [f16950, f574])).
fof(f16950, plain, (~ v8_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_118), inference(subsumption_resolution, [], [f16949, f575])).
fof(f16949, plain, (~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v8_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_118), inference(resolution, [], [f7354, f1103])).
fof(f1103, plain, ! [X0, X1] : (v17_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f438])).
fof(f438, plain, ! [X0] : (! [X1] : (((k2_waybel_1(X0, X0, X1) = k2_waybel34(X0, k2_yellow_2(X0, X0, X1), k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k1_waybel34(X0, k2_yellow_2(X0, X0, X1), k2_waybel_1(X0, X0, X1))) & v18_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v17_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f437])).
fof(f437, plain, ! [X0] : (! [X1] : (((k2_waybel_1(X0, X0, X1) = k2_waybel34(X0, k2_yellow_2(X0, X0, X1), k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k1_waybel34(X0, k2_yellow_2(X0, X0, X1), k2_waybel_1(X0, X0, X1))) & v18_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v17_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) | (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f150])).
fof(f150, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => ((k2_waybel_1(X0, X0, X1) = k2_waybel34(X0, k2_yellow_2(X0, X0, X1), k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k1_waybel34(X0, k2_yellow_2(X0, X0, X1), k2_waybel_1(X0, X0, X1))) & v18_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v17_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', t33_waybel34)).
fof(f7354, plain, (~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | spl50_118), inference(avatar_component_clause, [], [f7352])).
fof(f7352, plain, (spl50_118 <=> v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl50_118])])).
fof(f16943, plain, (spl50_873 | ~ spl50_110), inference(avatar_split_clause, [], [f16879, f7103, f16638])).
fof(f16638, plain, (spl50_873 <=> m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9))), introduced(avatar_definition, [new_symbols(naming, [spl50_873])])).
fof(f7103, plain, (spl50_110 <=> m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9))), introduced(avatar_definition, [new_symbols(naming, [spl50_110])])).
fof(f16879, plain, (m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ spl50_110), inference(resolution, [], [f7104, f1097])).
fof(f7104, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ spl50_110), inference(avatar_component_clause, [], [f7103])).
fof(f16928, plain, (~ spl50_873 | spl50_116 | spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109), inference(avatar_split_clause, [], [f16927, f7099, f2581, f2570, f2560, f2550, f2504, f2495, f1328, f7314, f16638])).
fof(f7314, plain, (spl50_116 <=> v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))), introduced(avatar_definition, [new_symbols(naming, [spl50_116])])).
fof(f2495, plain, (spl50_37 <=> l1_orders_2(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl50_37])])).
fof(f2504, plain, (spl50_39 <=> v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl50_39])])).
fof(f2550, plain, (spl50_48 <=> v3_orders_2(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl50_48])])).
fof(f2560, plain, (spl50_50 <=> v4_orders_2(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl50_50])])).
fof(f2570, plain, (spl50_52 <=> v2_lattice3(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl50_52])])).
fof(f2581, plain, (spl50_54 <=> v1_lattice3(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl50_54])])).
fof(f7099, plain, (spl50_109 <=> v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9))), introduced(avatar_definition, [new_symbols(naming, [spl50_109])])).
fof(f16927, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109)), inference(subsumption_resolution, [], [f16926, f565])).
fof(f16926, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109)), inference(subsumption_resolution, [], [f16925, f566])).
fof(f16925, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109)), inference(subsumption_resolution, [], [f16924, f567])).
fof(f16924, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109)), inference(subsumption_resolution, [], [f16923, f568])).
fof(f16923, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109)), inference(subsumption_resolution, [], [f16922, f569])).
fof(f16922, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109)), inference(subsumption_resolution, [], [f16921, f571])).
fof(f16921, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109)), inference(subsumption_resolution, [], [f16920, f2505])).
fof(f2505, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ spl50_39), inference(avatar_component_clause, [], [f2504])).
fof(f16920, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109)), inference(subsumption_resolution, [], [f16919, f2551])).
fof(f2551, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ spl50_48), inference(avatar_component_clause, [], [f2550])).
fof(f16919, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109)), inference(subsumption_resolution, [], [f16918, f2561])).
fof(f2561, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ spl50_50), inference(avatar_component_clause, [], [f2560])).
fof(f16918, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_52 | ~ spl50_54 | ~ spl50_109)), inference(subsumption_resolution, [], [f16917, f2582])).
fof(f2582, plain, (v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl50_54), inference(avatar_component_clause, [], [f2581])).
fof(f16917, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_52 | ~ spl50_109)), inference(subsumption_resolution, [], [f16916, f2571])).
fof(f2571, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl50_52), inference(avatar_component_clause, [], [f2570])).
fof(f16916, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_109)), inference(subsumption_resolution, [], [f16915, f2496])).
fof(f2496, plain, (l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ spl50_37), inference(avatar_component_clause, [], [f2495])).
fof(f16915, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_109)), inference(subsumption_resolution, [], [f16914, f2990])).
fof(f2990, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | spl50_7), inference(subsumption_resolution, [], [f2989, f1329])).
fof(f2989, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2988, f571])).
fof(f2988, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2987, f572])).
fof(f2987, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2986, f573])).
fof(f2986, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2967])).
fof(f2967, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f752, f1467])).
fof(f752, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v1_funct_1(k3_waybel_1(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f338])).
fof(f338, plain, ! [X0, X1, X2] : ((m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f337])).
fof(f337, plain, ! [X0, X1, X2] : ((m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f71])).
fof(f71, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', dt_k3_waybel_1)).
fof(f16914, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl50_109), inference(subsumption_resolution, [], [f7450, f7100])).
fof(f7100, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ spl50_109), inference(avatar_component_clause, [], [f7099])).
fof(f7450, plain, (v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(superposition, [], [f744, f4804])).
fof(f4804, plain, (k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))), inference(subsumption_resolution, [], [f4803, f565])).
fof(f4803, plain, ((k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4802, f566])).
fof(f4802, plain, ((k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4801, f567])).
fof(f4801, plain, ((k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4800, f568])).
fof(f4800, plain, ((k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4799, f569])).
fof(f4799, plain, ((k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4798, f570])).
fof(f4798, plain, ((k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4797, f571])).
fof(f4797, plain, ((k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4796, f572])).
fof(f4796, plain, ((k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4795, f573])).
fof(f4795, plain, ((k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4788, f574])).
fof(f4788, plain, ((k2_waybel_1(sK9, sK9, sK10) = k2_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k3_waybel_1(sK9, sK9, sK10))) | ~ v8_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(resolution, [], [f1106, f575])).
fof(f1106, plain, ! [X0, X1] : (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | (k2_waybel_1(X0, X0, X1) = k2_waybel34(X0, k2_yellow_2(X0, X0, X1), k3_waybel_1(X0, X0, X1))) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f438])).
fof(f744, plain, ! [X2, X0, X1] : (v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f332])).
fof(f332, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f331])).
fof(f331, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', dt_k2_waybel34)).
fof(f16860, plain, (spl50_7 | spl50_110), inference(avatar_contradiction_clause, [], [f16859])).
fof(f16859, plain, ($false | (spl50_7 | spl50_110)), inference(subsumption_resolution, [], [f16858, f1329])).
fof(f16858, plain, (v3_struct_0(sK9) | spl50_110), inference(subsumption_resolution, [], [f16857, f571])).
fof(f16857, plain, (~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_110), inference(subsumption_resolution, [], [f16856, f572])).
fof(f16856, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_110), inference(subsumption_resolution, [], [f16855, f573])).
fof(f16855, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_110), inference(subsumption_resolution, [], [f16854, f1467])).
fof(f16854, plain, (~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_110), inference(duplicate_literal_removal, [], [f16853])).
fof(f16853, plain, (~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_110), inference(resolution, [], [f7105, f754])).
fof(f754, plain, ! [X2, X0, X1] : (m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f338])).
fof(f7105, plain, (~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | spl50_110), inference(avatar_component_clause, [], [f7103])).
fof(f16670, plain, (~ spl50_109 | ~ spl50_873 | spl50_123 | spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54), inference(avatar_split_clause, [], [f16669, f2581, f2570, f2560, f2550, f2504, f2495, f1328, f7436, f16638, f7099])).
fof(f7436, plain, (spl50_123 <=> m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))), introduced(avatar_definition, [new_symbols(naming, [spl50_123])])).
fof(f16669, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54)), inference(subsumption_resolution, [], [f16668, f565])).
fof(f16668, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54)), inference(subsumption_resolution, [], [f16667, f566])).
fof(f16667, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54)), inference(subsumption_resolution, [], [f16666, f567])).
fof(f16666, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54)), inference(subsumption_resolution, [], [f16665, f568])).
fof(f16665, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54)), inference(subsumption_resolution, [], [f16664, f569])).
fof(f16664, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54)), inference(subsumption_resolution, [], [f16663, f571])).
fof(f16663, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54)), inference(subsumption_resolution, [], [f16662, f2505])).
fof(f16662, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54)), inference(subsumption_resolution, [], [f16661, f2551])).
fof(f16661, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54)), inference(subsumption_resolution, [], [f16660, f2561])).
fof(f16660, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_52 | ~ spl50_54)), inference(subsumption_resolution, [], [f16659, f2582])).
fof(f16659, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_52)), inference(subsumption_resolution, [], [f16658, f2571])).
fof(f16658, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37)), inference(subsumption_resolution, [], [f16657, f2496])).
fof(f16657, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_7), inference(subsumption_resolution, [], [f7449, f2990])).
fof(f7449, plain, (m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(superposition, [], [f745, f4804])).
fof(f745, plain, ! [X2, X0, X1] : (m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f332])).
fof(f16555, plain, (spl50_7 | spl50_109), inference(avatar_contradiction_clause, [], [f16554])).
fof(f16554, plain, ($false | (spl50_7 | spl50_109)), inference(subsumption_resolution, [], [f16553, f1329])).
fof(f16553, plain, (v3_struct_0(sK9) | spl50_109), inference(subsumption_resolution, [], [f16552, f571])).
fof(f16552, plain, (~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_109), inference(subsumption_resolution, [], [f16551, f572])).
fof(f16551, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_109), inference(subsumption_resolution, [], [f16550, f573])).
fof(f16550, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_109), inference(subsumption_resolution, [], [f16549, f1467])).
fof(f16549, plain, (~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_109), inference(duplicate_literal_removal, [], [f16548])).
fof(f16548, plain, (~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_109), inference(resolution, [], [f7101, f753])).
fof(f753, plain, ! [X2, X0, X1] : (v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f338])).
fof(f7101, plain, (~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | spl50_109), inference(avatar_component_clause, [], [f7099])).
fof(f7443, plain, (~ spl50_116 | ~ spl50_118 | ~ spl50_123 | ~ spl50_124 | spl50_111 | spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75), inference(avatar_split_clause, [], [f7434, f3348, f2581, f2570, f2560, f2550, f2504, f2495, f1328, f7107, f7440, f7436, f7352, f7314])).
fof(f7107, plain, (spl50_111 <=> v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9)), introduced(avatar_definition, [new_symbols(naming, [spl50_111])])).
fof(f3348, plain, (spl50_75 <=> v3_lattice3(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl50_75])])).
fof(f7434, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7433, f565])).
fof(f7433, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7432, f566])).
fof(f7432, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7431, f567])).
fof(f7431, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7430, f568])).
fof(f7430, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7429, f569])).
fof(f7429, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7428, f570])).
fof(f7428, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7427, f571])).
fof(f7427, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_39 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7426, f2505])).
fof(f7426, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_48 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7425, f2551])).
fof(f7425, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_50 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7424, f2561])).
fof(f7424, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_52 | ~ spl50_54 | ~ spl50_75)), inference(subsumption_resolution, [], [f7423, f2582])).
fof(f7423, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_52 | ~ spl50_75)), inference(subsumption_resolution, [], [f7422, f2571])).
fof(f7422, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37 | ~ spl50_75)), inference(subsumption_resolution, [], [f7421, f3349])).
fof(f3349, plain, (v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl50_75), inference(avatar_component_clause, [], [f3348])).
fof(f7421, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (spl50_7 | ~ spl50_37)), inference(subsumption_resolution, [], [f7420, f2496])).
fof(f7420, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_7), inference(subsumption_resolution, [], [f7298, f2913])).
fof(f2913, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | spl50_7), inference(subsumption_resolution, [], [f2912, f1329])).
fof(f2912, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2911, f571])).
fof(f2911, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2910, f572])).
fof(f2910, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2909, f573])).
fof(f2909, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2890])).
fof(f2890, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f746, f1467])).
fof(f746, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v1_funct_1(k2_waybel_1(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f334])).
fof(f334, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f333])).
fof(f333, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f68])).
fof(f68, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', dt_k2_waybel_1)).
fof(f7298, plain, (v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v22_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ m2_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v17_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(superposition, [], [f1101, f4784])).
fof(f4784, plain, (k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))), inference(subsumption_resolution, [], [f4783, f565])).
fof(f4783, plain, ((k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4782, f566])).
fof(f4782, plain, ((k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4781, f567])).
fof(f4781, plain, ((k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4780, f568])).
fof(f4780, plain, ((k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4779, f569])).
fof(f4779, plain, ((k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4778, f570])).
fof(f4778, plain, ((k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4777, f571])).
fof(f4777, plain, ((k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4776, f572])).
fof(f4776, plain, ((k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4775, f573])).
fof(f4775, plain, ((k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4768, f574])).
fof(f4768, plain, ((k3_waybel_1(sK9, sK9, sK10) = k1_waybel34(sK9, k2_yellow_2(sK9, sK9, sK10), k2_waybel_1(sK9, sK9, sK10))) | ~ v8_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(resolution, [], [f1105, f575])).
fof(f1105, plain, ! [X0, X1] : (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | (k3_waybel_1(X0, X0, X1) = k1_waybel34(X0, k2_yellow_2(X0, X0, X1), k2_waybel_1(X0, X0, X1))) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f438])).
fof(f1101, plain, ! [X2, X0, X1] : (v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ v22_waybel_0(X2, X0, X1) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f434])).
fof(f434, plain, ! [X0] : (! [X1] : (! [X2] : (v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ v22_waybel_0(X2, X0, X1) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f433])).
fof(f433, plain, ! [X0] : (! [X1] : (! [X2] : ((v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ v22_waybel_0(X2, X0, X1)) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f148])).
fof(f148, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => (v22_waybel_0(X2, X0, X1) => v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', t24_waybel34)).
fof(f7113, plain, (~ spl50_109 | ~ spl50_110 | ~ spl50_111 | spl50_112 | spl50_7 | spl50_26 | ~ spl50_37 | ~ spl50_39), inference(avatar_split_clause, [], [f7097, f2504, f2495, f1622, f1328, f7111, f7107, f7103, f7099])).
fof(f1622, plain, (spl50_26 <=> v3_struct_0(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl50_26])])).
fof(f7097, plain, (! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9))) | (spl50_7 | spl50_26 | ~ spl50_37 | ~ spl50_39)), inference(subsumption_resolution, [], [f7096, f1623])).
fof(f1623, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | spl50_26), inference(avatar_component_clause, [], [f1622])).
fof(f7096, plain, (! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10))) | (spl50_7 | ~ spl50_37 | ~ spl50_39)), inference(subsumption_resolution, [], [f7095, f2505])).
fof(f7095, plain, (! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10))) | (spl50_7 | ~ spl50_37)), inference(subsumption_resolution, [], [f7094, f2496])).
fof(f7094, plain, (! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10))) | spl50_7), inference(subsumption_resolution, [], [f7093, f1329])).
fof(f7093, plain, (! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | v3_struct_0(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10))) | spl50_7), inference(subsumption_resolution, [], [f7092, f565])).
fof(f7092, plain, (! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10))) | spl50_7), inference(subsumption_resolution, [], [f7091, f571])).
fof(f7091, plain, (! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ l1_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10))) | spl50_7), inference(subsumption_resolution, [], [f7090, f2990])).
fof(f7090, plain, (! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10))) | spl50_7), inference(subsumption_resolution, [], [f7087, f1131])).
fof(f7087, plain, (! [X0] : (r1_waybel_3(sK9, k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), X0), sK12) | ~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), X0, sK12) | ~ m1_subset_1(sK12, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ m1_subset_1(X0, u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_waybel34(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10))) | spl50_7), inference(superposition, [], [f733, f4744])).
fof(f4744, plain, ((sK12 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK12)) | spl50_7), inference(subsumption_resolution, [], [f4743, f1329])).
fof(f4743, plain, ((sK12 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK12)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f4742, f571])).
fof(f4742, plain, ((sK12 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK12)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f4741, f572])).
fof(f4741, plain, ((sK12 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK12)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f4740, f573])).
fof(f4740, plain, ((sK12 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK12)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f4739, f575])).
fof(f4739, plain, ((sK12 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK12)) | ~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f4732])).
fof(f4732, plain, ((sK12 = k7_yellow_2(u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), sK9, k3_waybel_1(sK9, sK9, sK10), sK12)) | ~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f1109, f1131])).
fof(f733, plain, ! [X6, X2, X0, X5, X1] : (r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, X5), k7_yellow_2(u1_struct_0(X0), X1, X2, X6)) | ~ r1_waybel_3(X0, X5, X6) | ~ m1_subset_1(X6, u1_struct_0(X0)) | ~ m1_subset_1(X5, u1_struct_0(X0)) | ~ v1_waybel34(X2, X0, X1) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f487])).
fof(f487, plain, ! [X0] : (! [X1] : (! [X2] : (((v1_waybel34(X2, X0, X1) | ((~ r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, sK15(X0, X1, X2)), k7_yellow_2(u1_struct_0(X0), X1, X2, sK16(X0, X1, X2))) & r1_waybel_3(X0, sK15(X0, X1, X2), sK16(X0, X1, X2)) & m1_subset_1(sK16(X0, X1, X2), u1_struct_0(X0))) & m1_subset_1(sK15(X0, X1, X2), u1_struct_0(X0)))) & (! [X5] : (! [X6] : (r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, X5), k7_yellow_2(u1_struct_0(X0), X1, X2, X6)) | ~ r1_waybel_3(X0, X5, X6) | ~ m1_subset_1(X6, u1_struct_0(X0))) | ~ m1_subset_1(X5, u1_struct_0(X0))) | ~ v1_waybel34(X2, X0, X1))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v2_orders_2(X1) | v3_struct_0(X1)) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16])], [f484, f486, f485])).
fof(f485, plain, ! [X2, X1, X0] : (? [X3] : (? [X4] : (~ r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, X3), k7_yellow_2(u1_struct_0(X0), X1, X2, X4)) & r1_waybel_3(X0, X3, X4) & m1_subset_1(X4, u1_struct_0(X0))) & m1_subset_1(X3, u1_struct_0(X0))) => (? [X4] : (~ r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, sK15(X0, X1, X2)), k7_yellow_2(u1_struct_0(X0), X1, X2, X4)) & r1_waybel_3(X0, sK15(X0, X1, X2), X4) & m1_subset_1(X4, u1_struct_0(X0))) & m1_subset_1(sK15(X0, X1, X2), u1_struct_0(X0)))), introduced(choice_axiom, [])).
fof(f486, plain, ! [X2, X1, X0] : (? [X4] : (~ r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, sK15(X0, X1, X2)), k7_yellow_2(u1_struct_0(X0), X1, X2, X4)) & r1_waybel_3(X0, sK15(X0, X1, X2), X4) & m1_subset_1(X4, u1_struct_0(X0))) => (~ r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, sK15(X0, X1, X2)), k7_yellow_2(u1_struct_0(X0), X1, X2, sK16(X0, X1, X2))) & r1_waybel_3(X0, sK15(X0, X1, X2), sK16(X0, X1, X2)) & m1_subset_1(sK16(X0, X1, X2), u1_struct_0(X0)))), introduced(choice_axiom, [])).
fof(f484, plain, ! [X0] : (! [X1] : (! [X2] : (((v1_waybel34(X2, X0, X1) | ? [X3] : (? [X4] : (~ r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, X3), k7_yellow_2(u1_struct_0(X0), X1, X2, X4)) & r1_waybel_3(X0, X3, X4) & m1_subset_1(X4, u1_struct_0(X0))) & m1_subset_1(X3, u1_struct_0(X0)))) & (! [X5] : (! [X6] : (r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, X5), k7_yellow_2(u1_struct_0(X0), X1, X2, X6)) | ~ r1_waybel_3(X0, X5, X6) | ~ m1_subset_1(X6, u1_struct_0(X0))) | ~ m1_subset_1(X5, u1_struct_0(X0))) | ~ v1_waybel34(X2, X0, X1))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v2_orders_2(X1) | v3_struct_0(X1)) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(rectify, [], [f483])).
fof(f483, plain, ! [X0] : (! [X1] : (! [X2] : (((v1_waybel34(X2, X0, X1) | ? [X3] : (? [X4] : (~ r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, X3), k7_yellow_2(u1_struct_0(X0), X1, X2, X4)) & r1_waybel_3(X0, X3, X4) & m1_subset_1(X4, u1_struct_0(X0))) & m1_subset_1(X3, u1_struct_0(X0)))) & (! [X3] : (! [X4] : (r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, X3), k7_yellow_2(u1_struct_0(X0), X1, X2, X4)) | ~ r1_waybel_3(X0, X3, X4) | ~ m1_subset_1(X4, u1_struct_0(X0))) | ~ m1_subset_1(X3, u1_struct_0(X0))) | ~ v1_waybel34(X2, X0, X1))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v2_orders_2(X1) | v3_struct_0(X1)) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f327])).
fof(f327, plain, ! [X0] : (! [X1] : (! [X2] : ((v1_waybel34(X2, X0, X1) <=> ! [X3] : (! [X4] : (r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, X3), k7_yellow_2(u1_struct_0(X0), X1, X2, X4)) | ~ r1_waybel_3(X0, X3, X4) | ~ m1_subset_1(X4, u1_struct_0(X0))) | ~ m1_subset_1(X3, u1_struct_0(X0)))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v2_orders_2(X1) | v3_struct_0(X1)) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f326])).
fof(f326, plain, ! [X0] : (! [X1] : (! [X2] : ((v1_waybel34(X2, X0, X1) <=> ! [X3] : (! [X4] : ((r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, X3), k7_yellow_2(u1_struct_0(X0), X1, X2, X4)) | ~ r1_waybel_3(X0, X3, X4)) | ~ m1_subset_1(X4, u1_struct_0(X0))) | ~ m1_subset_1(X3, u1_struct_0(X0)))) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v2_orders_2(X1) | v3_struct_0(X1))) | (~ l1_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((l1_orders_2(X1) & v2_orders_2(X1) & ~ v3_struct_0(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => (v1_waybel34(X2, X0, X1) <=> ! [X3] : (m1_subset_1(X3, u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4, u1_struct_0(X0)) => (r1_waybel_3(X0, X3, X4) => r1_waybel_3(X1, k7_yellow_2(u1_struct_0(X0), X1, X2, X3), k7_yellow_2(u1_struct_0(X0), X1, X2, X4))))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', d8_waybel34)).
fof(f4763, plain, (spl50_73 | spl50_7 | ~ spl50_72), inference(avatar_split_clause, [], [f4762, f3246, f1328, f3250])).
fof(f3250, plain, (spl50_73 <=> v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9)), introduced(avatar_definition, [new_symbols(naming, [spl50_73])])).
fof(f4762, plain, (v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | (spl50_7 | ~ spl50_72)), inference(subsumption_resolution, [], [f4761, f1329])).
fof(f4761, plain, (v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(sK9) | (spl50_7 | ~ spl50_72)), inference(subsumption_resolution, [], [f4760, f571])).
fof(f4760, plain, (v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | (spl50_7 | ~ spl50_72)), inference(subsumption_resolution, [], [f4754, f3247])).
fof(f4754, plain, (v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl50_7), inference(resolution, [], [f3968, f591])).
fof(f591, plain, ! [X0, X1] : (~ v8_yellow_0(X1, X0) | v6_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f245])).
fof(f245, plain, ! [X0] : (! [X1] : (v6_yellow_0(X1, X0) | ~ v8_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f244])).
fof(f244, plain, ! [X0] : (! [X1] : ((v6_yellow_0(X1, X0) | ~ v8_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v8_yellow_0(X1, X0) => v6_yellow_0(X1, X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', cc10_yellow_0)).
fof(f4210, plain, (~ spl50_70 | spl50_75), inference(avatar_contradiction_clause, [], [f4209])).
fof(f4209, plain, ($false | (~ spl50_70 | spl50_75)), inference(subsumption_resolution, [], [f4208, f565])).
fof(f4208, plain, (~ v2_orders_2(sK9) | (~ spl50_70 | spl50_75)), inference(subsumption_resolution, [], [f4207, f566])).
fof(f4207, plain, (~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl50_70 | spl50_75)), inference(subsumption_resolution, [], [f4206, f567])).
fof(f4206, plain, (~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl50_70 | spl50_75)), inference(subsumption_resolution, [], [f4205, f568])).
fof(f4205, plain, (~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl50_70 | spl50_75)), inference(subsumption_resolution, [], [f4204, f569])).
fof(f4204, plain, (~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl50_70 | spl50_75)), inference(subsumption_resolution, [], [f4203, f570])).
fof(f4203, plain, (~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl50_70 | spl50_75)), inference(subsumption_resolution, [], [f4202, f571])).
fof(f4202, plain, (~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl50_70 | spl50_75)), inference(subsumption_resolution, [], [f4201, f572])).
fof(f4201, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl50_70 | spl50_75)), inference(subsumption_resolution, [], [f4200, f573])).
fof(f4200, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl50_70 | spl50_75)), inference(subsumption_resolution, [], [f4199, f2714])).
fof(f2714, plain, (v6_waybel_1(sK10, sK9) | ~ spl50_70), inference(avatar_component_clause, [], [f2713])).
fof(f2713, plain, (spl50_70 <=> v6_waybel_1(sK10, sK9)), introduced(avatar_definition, [new_symbols(naming, [spl50_70])])).
fof(f4199, plain, (~ v6_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_75), inference(subsumption_resolution, [], [f4181, f3535])).
fof(f3535, plain, (~ sP2(sK10, sK9) | spl50_75), inference(resolution, [], [f3350, f791])).
fof(f791, plain, ! [X0, X1] : (v3_lattice3(k2_yellow_2(X1, X1, X0)) | ~ sP2(X0, X1)), inference(cnf_transformation, [], [f501])).
fof(f501, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X1, X1, X0)) & v2_lattice3(k2_yellow_2(X1, X1, X0)) & v1_lattice3(k2_yellow_2(X1, X1, X0)) & v25_waybel_0(k2_yellow_2(X1, X1, X0)) & v24_waybel_0(k2_yellow_2(X1, X1, X0)) & v4_yellow_0(k2_yellow_2(X1, X1, X0), X1) & v3_yellow_0(k2_yellow_2(X1, X1, X0)) & v2_yellow_0(k2_yellow_2(X1, X1, X0)) & v1_yellow_0(k2_yellow_2(X1, X1, X0)) & v4_orders_2(k2_yellow_2(X1, X1, X0)) & v3_orders_2(k2_yellow_2(X1, X1, X0)) & v2_orders_2(k2_yellow_2(X1, X1, X0)) & v1_orders_2(k2_yellow_2(X1, X1, X0)) & ~ v3_struct_0(k2_yellow_2(X1, X1, X0))) | ~ sP2(X0, X1)), inference(rectify, [], [f500])).
fof(f500, plain, ! [X1, X0] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ sP2(X1, X0)), inference(nnf_transformation, [], [f458])).
fof(f458, plain, ! [X1, X0] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ sP2(X1, X0)), inference(usedef, [], [e458])).
fof(e458, plain, ! [X1, X0] : (sP2(X1, X0) <=> (v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f3350, plain, (~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | spl50_75), inference(avatar_component_clause, [], [f3348])).
fof(f4181, plain, (sP2(sK10, sK9) | ~ v6_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(resolution, [], [f792, f1467])).
fof(f792, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | sP2(X1, X0) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f459])).
fof(f459, plain, ! [X0, X1] : (sP2(X1, X0) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(definition_folding, [], [f351, e458])).
fof(f351, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f350])).
fof(f350, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f90])).
fof(f90, plain, ! [X0, X1] : ((m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', fc10_waybel34)).
fof(f3362, plain, (spl50_52 | ~ spl50_75 | spl50_26 | ~ spl50_37), inference(avatar_split_clause, [], [f3361, f2495, f1622, f3348, f2570])).
fof(f3361, plain, (~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | (spl50_26 | ~ spl50_37)), inference(subsumption_resolution, [], [f3338, f1623])).
fof(f3338, plain, (~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl50_37), inference(resolution, [], [f2496, f665])).
fof(f665, plain, ! [X0] : (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | v3_struct_0(X0) | v2_lattice3(X0)), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ! [X0] : ((v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)) | ~ v3_lattice3(X0) | v3_struct_0(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f277])).
fof(f277, plain, ! [X0] : (((v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)) | (~ v3_lattice3(X0) | v3_struct_0(X0))) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f33])).
fof(f33, plain, ! [X0] : (l1_orders_2(X0) => ((v3_lattice3(X0) & ~ v3_struct_0(X0)) => (v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', cc1_yellow_0)).
fof(f3328, plain, (spl50_37 | ~ spl50_72), inference(avatar_split_clause, [], [f3327, f3246, f2495])).
fof(f3327, plain, (l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ spl50_72), inference(subsumption_resolution, [], [f3323, f571])).
fof(f3323, plain, (l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ spl50_72), inference(resolution, [], [f3247, f761])).
fof(f761, plain, ! [X0, X1] : (~ m1_yellow_0(X1, X0) | l1_orders_2(X1) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f344])).
fof(f344, plain, ! [X0] : (! [X1] : (l1_orders_2(X1) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f79])).
fof(f79, plain, ! [X0] : (l1_orders_2(X0) => ! [X1] : (m1_yellow_0(X1, X0) => l1_orders_2(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', dt_m1_yellow_0)).
fof(f3319, plain, (spl50_72 | spl50_7), inference(avatar_split_clause, [], [f3318, f1328, f3246])).
fof(f3318, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl50_7), inference(subsumption_resolution, [], [f3317, f1329])).
fof(f3317, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3316, f571])).
fof(f3316, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3315, f572])).
fof(f3315, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3294, f573])).
fof(f3294, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f3275])).
fof(f3275, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f751, f1467])).
fof(f751, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f336])).
fof(f3274, plain, (~ spl50_72 | spl50_39 | spl50_7), inference(avatar_split_clause, [], [f3273, f1328, f2504, f3246])).
fof(f3273, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl50_7), inference(subsumption_resolution, [], [f3272, f565])).
fof(f3272, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v2_orders_2(sK9) | spl50_7), inference(subsumption_resolution, [], [f3239, f571])).
fof(f3239, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v2_orders_2(sK9) | spl50_7), inference(resolution, [], [f3219, f714])).
fof(f714, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v2_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f309])).
fof(f309, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v2_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f308])).
fof(f308, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v2_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f52])).
fof(f52, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v2_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', cc6_yellow_0)).
fof(f3271, plain, (~ spl50_72 | spl50_48 | spl50_7), inference(avatar_split_clause, [], [f3270, f1328, f2550, f3246])).
fof(f3270, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl50_7), inference(subsumption_resolution, [], [f3269, f566])).
fof(f3269, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v3_orders_2(sK9) | spl50_7), inference(subsumption_resolution, [], [f3238, f571])).
fof(f3238, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v3_orders_2(sK9) | spl50_7), inference(resolution, [], [f3219, f719])).
fof(f719, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v3_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v3_orders_2(X0)), inference(cnf_transformation, [], [f313])).
fof(f313, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v3_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v3_orders_2(X0)), inference(flattening, [], [f312])).
fof(f312, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v3_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v3_orders_2(X0))), inference(ennf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : ((l1_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v3_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', cc7_yellow_0)).
fof(f3268, plain, (~ spl50_72 | spl50_50 | spl50_7), inference(avatar_split_clause, [], [f3267, f1328, f2560, f3246])).
fof(f3267, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl50_7), inference(subsumption_resolution, [], [f3266, f567])).
fof(f3266, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v4_orders_2(sK9) | spl50_7), inference(subsumption_resolution, [], [f3237, f571])).
fof(f3237, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | spl50_7), inference(resolution, [], [f3219, f724])).
fof(f724, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v4_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0)), inference(cnf_transformation, [], [f317])).
fof(f317, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v4_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0)), inference(flattening, [], [f316])).
fof(f316, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v4_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v4_orders_2(X0))), inference(ennf_transformation, [], [f56])).
fof(f56, plain, ! [X0] : ((l1_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v4_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', cc8_yellow_0)).
fof(f3253, plain, (~ spl50_72 | spl50_54 | ~ spl50_73 | spl50_7 | spl50_26), inference(avatar_split_clause, [], [f3244, f1622, f1328, f3250, f2581, f3246])).
fof(f3244, plain, (~ v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | (spl50_7 | spl50_26)), inference(subsumption_resolution, [], [f3243, f566])).
fof(f3243, plain, (~ v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v3_orders_2(sK9) | (spl50_7 | spl50_26)), inference(subsumption_resolution, [], [f3242, f567])).
fof(f3242, plain, (~ v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | (spl50_7 | spl50_26)), inference(subsumption_resolution, [], [f3241, f568])).
fof(f3241, plain, (~ v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | (spl50_7 | spl50_26)), inference(subsumption_resolution, [], [f3240, f571])).
fof(f3240, plain, (~ v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | (spl50_7 | spl50_26)), inference(subsumption_resolution, [], [f3234, f1623])).
fof(f3234, plain, (~ v6_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl50_7), inference(resolution, [], [f3219, f615])).
fof(f615, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | ~ v6_yellow_0(X1, X0) | v1_lattice3(X1) | v3_struct_0(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0)), inference(cnf_transformation, [], [f253])).
fof(f253, plain, ! [X0] : (! [X1] : ((v6_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1)) | ~ v6_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0)), inference(flattening, [], [f252])).
fof(f252, plain, ! [X0] : (! [X1] : (((v6_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1)) | (~ v6_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1))) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0] : ((l1_orders_2(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => ((v6_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & ~ v3_struct_0(X1)) => (v6_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', cc12_yellow_0)).
fof(f3029, plain, (~ spl50_26 | spl50_7), inference(avatar_split_clause, [], [f3028, f1328, f1622])).
fof(f3028, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | spl50_7), inference(subsumption_resolution, [], [f3027, f1329])).
fof(f3027, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3026, f571])).
fof(f3026, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3025, f572])).
fof(f3025, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3024, f573])).
fof(f3024, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f3005])).
fof(f3005, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f834, f1467])).
fof(f834, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v3_struct_0(k2_yellow_2(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f368])).
fof(f368, plain, ! [X0, X1, X2] : ((v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f367])).
fof(f367, plain, ! [X0, X1, X2] : ((v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f99])).
fof(f99, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', fc2_yellow_2)).
fof(f2846, plain, (spl50_70 | spl50_7), inference(avatar_split_clause, [], [f2845, f1328, f2713])).
fof(f2845, plain, (v6_waybel_1(sK10, sK9) | spl50_7), inference(subsumption_resolution, [], [f2844, f1329])).
fof(f2844, plain, (v6_waybel_1(sK10, sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2843, f571])).
fof(f2843, plain, (v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2842, f572])).
fof(f2842, plain, (~ v1_funct_1(sK10) | v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2841, f573])).
fof(f2841, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2824, f574])).
fof(f2824, plain, (~ v8_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f707, f1467])).
fof(f707, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | v6_waybel_1(X1, X0) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f301])).
fof(f301, plain, ! [X0] : (! [X1] : ((v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1)) | ~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0))) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f300])).
fof(f300, plain, ! [X0] : (! [X1] : (((v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1)) | (~ v8_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0))) | (~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f216])).
fof(f216, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) => ((v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1))))), inference(pure_predicate_removal, [], [f48])).
fof(f48, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) => ((v8_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v6_waybel_1(X1, X0) & v1_partfun1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', cc4_waybel_1)).
fof(f1411, plain, ~ spl50_7, inference(avatar_contradiction_clause, [], [f1410])).
fof(f1410, plain, ($false | ~ spl50_7), inference(subsumption_resolution, [], [f1409, f571])).
fof(f1409, plain, (~ l1_orders_2(sK9) | ~ spl50_7), inference(subsumption_resolution, [], [f1405, f568])).
fof(f1405, plain, (~ v1_lattice3(sK9) | ~ l1_orders_2(sK9) | ~ spl50_7), inference(resolution, [], [f1330, f644])).
fof(f644, plain, ! [X0] : (~ v3_struct_0(X0) | ~ v1_lattice3(X0) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f266])).
fof(f266, plain, ! [X0] : (~ v3_struct_0(X0) | ~ v1_lattice3(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f265])).
fof(f265, plain, ! [X0] : ((~ v3_struct_0(X0) | ~ v1_lattice3(X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ! [X0] : (l1_orders_2(X0) => (v1_lattice3(X0) => ~ v3_struct_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT366+1.p', cc1_lattice3)).
fof(f1330, plain, (v3_struct_0(sK9) | ~ spl50_7), inference(avatar_component_clause, [], [f1328])).
fof(f1130, plain, (spl50_1 | spl50_2), inference(avatar_split_clause, [], [f583, f1126, f1122])).
fof(f583, plain, (r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK14) | r1_waybel_3(sK9, sK11, sK12)), inference(cnf_transformation, [], [f480])).
fof(f1129, plain, (~ spl50_1 | ~ spl50_2), inference(avatar_split_clause, [], [f584, f1126, f1122])).
fof(f584, plain, (~ r1_waybel_3(k2_yellow_2(sK9, sK9, sK10), sK13, sK14) | ~ r1_waybel_3(sK9, sK11, sK12)), inference(cnf_transformation, [], [f480])).