fof(f3254, plain, $false, inference(avatar_sat_refutation, [], [f942, f943, f944, f945, f946, f2960, f2980, f2997, f3042, f3062, f3084, f3103, f3125, f3144, f3165, f3189, f3253])).
fof(f3253, plain, (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70), inference(avatar_contradiction_clause, [], [f3252])).
fof(f3252, plain, ($false | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3251, f466])).
fof(f466, plain, v2_orders_2(sK7), inference(cnf_transformation, [], [f388])).
fof(f388, plain, ((((~ m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8)))) | ~ v22_waybel_0(sK9, sK7, sK8)) & ((m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) & v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) & v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) & v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))))) | v22_waybel_0(sK9, sK7, sK8)) & m2_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) & v17_waybel_0(sK9, sK7, sK8) & v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) & v1_funct_1(sK9)) & v2_waybel_8(sK8) & l1_orders_2(sK8) & v3_lattice3(sK8) & v2_lattice3(sK8) & v1_lattice3(sK8) & v4_orders_2(sK8) & v3_orders_2(sK8) & v2_orders_2(sK8)) & l1_orders_2(sK7) & v3_lattice3(sK7) & v2_lattice3(sK7) & v1_lattice3(sK7) & v4_orders_2(sK7) & v3_orders_2(sK7) & v2_orders_2(sK7)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8, sK9])], [f384, f387, f386, f385])).
fof(f385, plain, (? [X0] : (? [X1] : (? [X2] : ((~ m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1)))) | ~ v22_waybel_0(X2, X0, X1)) & ((m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))))) | v22_waybel_0(X2, X0, X1)) & m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) & v2_waybel_8(X1) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (? [X1] : (? [X2] : ((~ m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(sK7))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(sK7)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1)))) | ~ v22_waybel_0(X2, sK7, X1)) & ((m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(sK7))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(sK7)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(sK7))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))))) | v22_waybel_0(X2, sK7, X1)) & m2_relset_1(X2, u1_struct_0(sK7), u1_struct_0(X1)) & v17_waybel_0(X2, sK7, X1) & v1_funct_2(X2, u1_struct_0(sK7), u1_struct_0(X1)) & v1_funct_1(X2)) & v2_waybel_8(X1) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) & l1_orders_2(sK7) & v3_lattice3(sK7) & v2_lattice3(sK7) & v1_lattice3(sK7) & v4_orders_2(sK7) & v3_orders_2(sK7) & v2_orders_2(sK7))), introduced(choice_axiom, [])).
fof(f386, plain, (? [X1] : (? [X2] : ((~ m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(sK7))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(sK7)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1)))) | ~ v22_waybel_0(X2, sK7, X1)) & ((m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(sK7))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(sK7)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(sK7))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(sK7), k1_waybel34(sK7, X1, X2), u1_struct_0(k1_waybel_8(X1))))) | v22_waybel_0(X2, sK7, X1)) & m2_relset_1(X2, u1_struct_0(sK7), u1_struct_0(X1)) & v17_waybel_0(X2, sK7, X1) & v1_funct_2(X2, u1_struct_0(sK7), u1_struct_0(X1)) & v1_funct_1(X2)) & v2_waybel_8(X1) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => (? [X2] : ((~ m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8)))) | ~ v22_waybel_0(X2, sK7, sK8)) & ((m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) & v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) & v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) & v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))))) | v22_waybel_0(X2, sK7, sK8)) & m2_relset_1(X2, u1_struct_0(sK7), u1_struct_0(sK8)) & v17_waybel_0(X2, sK7, sK8) & v1_funct_2(X2, u1_struct_0(sK7), u1_struct_0(sK8)) & v1_funct_1(X2)) & v2_waybel_8(sK8) & l1_orders_2(sK8) & v3_lattice3(sK8) & v2_lattice3(sK8) & v1_lattice3(sK8) & v4_orders_2(sK8) & v3_orders_2(sK8) & v2_orders_2(sK8))), introduced(choice_axiom, [])).
fof(f387, plain, (? [X2] : ((~ m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8)))) | ~ v22_waybel_0(X2, sK7, sK8)) & ((m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) & v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) & v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) & v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, X2), u1_struct_0(k1_waybel_8(sK8))))) | v22_waybel_0(X2, sK7, sK8)) & m2_relset_1(X2, u1_struct_0(sK7), u1_struct_0(sK8)) & v17_waybel_0(X2, sK7, sK8) & v1_funct_2(X2, u1_struct_0(sK7), u1_struct_0(sK8)) & v1_funct_1(X2)) => ((~ m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8)))) | ~ v22_waybel_0(sK9, sK7, sK8)) & ((m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) & v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) & v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) & v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))))) | v22_waybel_0(sK9, sK7, sK8)) & m2_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) & v17_waybel_0(sK9, sK7, sK8) & v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) & v1_funct_1(sK9))), introduced(choice_axiom, [])).
fof(f384, plain, ? [X0] : (? [X1] : (? [X2] : ((~ m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1)))) | ~ v22_waybel_0(X2, X0, X1)) & ((m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))))) | v22_waybel_0(X2, X0, X1)) & m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) & v2_waybel_8(X1) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(flattening, [], [f383])).
fof(f383, plain, ? [X0] : (? [X1] : (? [X2] : ((((~ m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))))) | ~ v22_waybel_0(X2, X0, X1)) & ((m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))))) | v22_waybel_0(X2, X0, X1))) & m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) & v2_waybel_8(X1) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(nnf_transformation, [], [f191])).
fof(f191, plain, ? [X0] : (? [X1] : (? [X2] : (~ (v22_waybel_0(X2, X0, X1) <=> (m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1)))))) & m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) & v2_waybel_8(X1) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(flattening, [], [f190])).
fof(f190, plain, ? [X0] : (? [X1] : ((? [X2] : (~ (v22_waybel_0(X2, X0, X1) <=> (m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1)))))) & (m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2))) & v2_waybel_8(X1)) & (l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1))) & (l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => (v2_waybel_8(X1) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => (v22_waybel_0(X2, X0, X1) <=> (m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1)))))))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => (v2_waybel_8(X1) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => (v22_waybel_0(X2, X0, X1) <=> (m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), k1_waybel34(X0, X1, X2), u1_struct_0(k1_waybel_8(X1)))))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT380+1.p', t71_waybel34)).
fof(f3251, plain, (~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3250, f467])).
fof(f467, plain, v3_orders_2(sK7), inference(cnf_transformation, [], [f388])).
fof(f3250, plain, (~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3249, f468])).
fof(f468, plain, v4_orders_2(sK7), inference(cnf_transformation, [], [f388])).
fof(f3249, plain, (~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3248, f469])).
fof(f469, plain, v1_lattice3(sK7), inference(cnf_transformation, [], [f388])).
fof(f3248, plain, (~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3247, f470])).
fof(f470, plain, v2_lattice3(sK7), inference(cnf_transformation, [], [f388])).
fof(f3247, plain, (~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3246, f471])).
fof(f471, plain, v3_lattice3(sK7), inference(cnf_transformation, [], [f388])).
fof(f3246, plain, (~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3245, f472])).
fof(f472, plain, l1_orders_2(sK7), inference(cnf_transformation, [], [f388])).
fof(f3245, plain, (~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3244, f473])).
fof(f473, plain, v2_orders_2(sK8), inference(cnf_transformation, [], [f388])).
fof(f3244, plain, (~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3243, f474])).
fof(f474, plain, v3_orders_2(sK8), inference(cnf_transformation, [], [f388])).
fof(f3243, plain, (~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3242, f475])).
fof(f475, plain, v4_orders_2(sK8), inference(cnf_transformation, [], [f388])).
fof(f3242, plain, (~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3241, f476])).
fof(f476, plain, v1_lattice3(sK8), inference(cnf_transformation, [], [f388])).
fof(f3241, plain, (~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3240, f477])).
fof(f477, plain, v2_lattice3(sK8), inference(cnf_transformation, [], [f388])).
fof(f3240, plain, (~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3239, f478])).
fof(f478, plain, v3_lattice3(sK8), inference(cnf_transformation, [], [f388])).
fof(f3239, plain, (~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3238, f479])).
fof(f479, plain, l1_orders_2(sK8), inference(cnf_transformation, [], [f388])).
fof(f3238, plain, (~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3237, f2473])).
fof(f2473, plain, v1_funct_1(k1_waybel34(sK7, sK8, sK9)), inference(subsumption_resolution, [], [f2472, f466])).
fof(f2472, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2471, f467])).
fof(f2471, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2470, f468])).
fof(f2470, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2469, f469])).
fof(f2469, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2468, f470])).
fof(f2468, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2467, f472])).
fof(f2467, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2466, f473])).
fof(f2466, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2465, f474])).
fof(f2465, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2464, f475])).
fof(f2464, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2463, f476])).
fof(f2463, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2462, f477])).
fof(f2462, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2461, f479])).
fof(f2461, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2460, f481])).
fof(f481, plain, v1_funct_1(sK9), inference(cnf_transformation, [], [f388])).
fof(f2460, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(subsumption_resolution, [], [f2441, f482])).
fof(f482, plain, v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)), inference(cnf_transformation, [], [f388])).
fof(f2441, plain, (v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7)), inference(resolution, [], [f655, f1185])).
fof(f1185, plain, m1_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)), inference(resolution, [], [f887, f484])).
fof(f484, plain, m2_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)), inference(cnf_transformation, [], [f388])).
fof(f887, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f454])).
fof(f454, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f122])).
fof(f122, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT380+1.p', redefinition_m2_relset_1)).
fof(f655, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v1_funct_1(k1_waybel34(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f277])).
fof(f277, plain, ! [X0, X1, X2] : ((m2_relset_1(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(k1_waybel34(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f276])).
fof(f276, plain, ! [X0, X1, X2] : ((m2_relset_1(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(k1_waybel34(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f58])).
fof(f58, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (m2_relset_1(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(k1_waybel34(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT380+1.p', dt_k1_waybel34)).
fof(f3237, plain, (~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3236, f2927])).
fof(f2927, plain, (v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ spl43_67), inference(avatar_component_clause, [], [f2926])).
fof(f2926, plain, (spl43_67 <=> v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7))), introduced(avatar_definition, [new_symbols(naming, [spl43_67])])).
fof(f3236, plain, (~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3235, f2931])).
fof(f2931, plain, (v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ spl43_68), inference(avatar_component_clause, [], [f2930])).
fof(f2930, plain, (spl43_68 <=> v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7)), introduced(avatar_definition, [new_symbols(naming, [spl43_68])])).
fof(f3235, plain, (~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3234, f2935])).
fof(f2935, plain, (m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ spl43_69), inference(avatar_component_clause, [], [f2934])).
fof(f2934, plain, (spl43_69 <=> m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7))), introduced(avatar_definition, [new_symbols(naming, [spl43_69])])).
fof(f3234, plain, (~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_2 | ~ spl43_70)), inference(subsumption_resolution, [], [f3233, f2940])).
fof(f2940, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ spl43_70), inference(avatar_component_clause, [], [f2938])).
fof(f2938, plain, (spl43_70 <=> v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7)), introduced(avatar_definition, [new_symbols(naming, [spl43_70])])).
fof(f3233, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_2), inference(resolution, [], [f929, f903])).
fof(f903, plain, ! [X2, X0, X1] : (v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1)))) | ~ v3_waybel34(X2, X1, X0) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f457])).
fof(f457, plain, ! [X0] : (! [X1] : (! [X2] : (((v3_waybel34(X2, X1, X0) | ~ m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))))) & ((m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))))) | ~ v3_waybel34(X2, X1, X0))) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f456])).
fof(f456, plain, ! [X0] : (! [X1] : (! [X2] : (((v3_waybel34(X2, X1, X0) | (~ m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1)))))) & ((m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))))) | ~ v3_waybel34(X2, X1, X0))) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(nnf_transformation, [], [f365])).
fof(f365, plain, ! [X0] : (! [X1] : (! [X2] : ((v3_waybel34(X2, X1, X0) <=> (m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1)))))) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f364])).
fof(f364, plain, ! [X0] : (! [X1] : (! [X2] : ((v3_waybel34(X2, X1, X0) <=> (m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1)))))) | (~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f136])).
fof(f136, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2)) => (v3_waybel34(X2, X1, X0) <=> (m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) & v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) & v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT380+1.p', t70_waybel34)).
fof(f929, plain, (~ v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8)))) | spl43_2), inference(avatar_component_clause, [], [f927])).
fof(f927, plain, (spl43_2 <=> v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))))), introduced(avatar_definition, [new_symbols(naming, [spl43_2])])).
fof(f3189, plain, (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70), inference(avatar_contradiction_clause, [], [f3188])).
fof(f3188, plain, ($false | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3187, f466])).
fof(f3187, plain, (~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3186, f467])).
fof(f3186, plain, (~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3185, f468])).
fof(f3185, plain, (~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3184, f469])).
fof(f3184, plain, (~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3183, f470])).
fof(f3183, plain, (~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3182, f471])).
fof(f3182, plain, (~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3181, f472])).
fof(f3181, plain, (~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3180, f473])).
fof(f3180, plain, (~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3179, f474])).
fof(f3179, plain, (~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3178, f475])).
fof(f3178, plain, (~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3177, f476])).
fof(f3177, plain, (~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3176, f477])).
fof(f3176, plain, (~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3175, f478])).
fof(f3175, plain, (~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3174, f479])).
fof(f3174, plain, (~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3173, f2473])).
fof(f3173, plain, (~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3172, f2927])).
fof(f3172, plain, (~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3171, f2931])).
fof(f3171, plain, (~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3170, f2935])).
fof(f3170, plain, (~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_5 | ~ spl43_70)), inference(subsumption_resolution, [], [f3169, f2940])).
fof(f3169, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_5), inference(resolution, [], [f941, f906])).
fof(f906, plain, ! [X2, X0, X1] : (m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v3_waybel34(X2, X1, X0) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f457])).
fof(f941, plain, (~ m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | spl43_5), inference(avatar_component_clause, [], [f939])).
fof(f939, plain, (spl43_5 <=> m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7)))), introduced(avatar_definition, [new_symbols(naming, [spl43_5])])).
fof(f3165, plain, (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70), inference(avatar_contradiction_clause, [], [f3164])).
fof(f3164, plain, ($false | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3163, f466])).
fof(f3163, plain, (~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3162, f467])).
fof(f3162, plain, (~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3161, f468])).
fof(f3161, plain, (~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3160, f469])).
fof(f3160, plain, (~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3159, f470])).
fof(f3159, plain, (~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3158, f471])).
fof(f3158, plain, (~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3157, f472])).
fof(f3157, plain, (~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3156, f473])).
fof(f3156, plain, (~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3155, f474])).
fof(f3155, plain, (~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3154, f475])).
fof(f3154, plain, (~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3153, f476])).
fof(f3153, plain, (~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3152, f477])).
fof(f3152, plain, (~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3151, f478])).
fof(f3151, plain, (~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3150, f479])).
fof(f3150, plain, (~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3149, f2473])).
fof(f3149, plain, (~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3148, f2927])).
fof(f3148, plain, (~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3147, f2931])).
fof(f3147, plain, (~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3146, f2935])).
fof(f3146, plain, (~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_4 | ~ spl43_70)), inference(subsumption_resolution, [], [f3145, f2940])).
fof(f3145, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_4), inference(resolution, [], [f937, f905])).
fof(f905, plain, ! [X2, X0, X1] : (v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) | ~ v3_waybel34(X2, X1, X0) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f457])).
fof(f937, plain, (~ v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) | spl43_4), inference(avatar_component_clause, [], [f935])).
fof(f935, plain, (spl43_4 <=> v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7))), introduced(avatar_definition, [new_symbols(naming, [spl43_4])])).
fof(f3144, plain, (~ spl43_70 | spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69), inference(avatar_split_clause, [], [f3143, f2934, f2930, f2926, f931, f2938])).
fof(f931, plain, (spl43_3 <=> v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7)))), introduced(avatar_definition, [new_symbols(naming, [spl43_3])])).
fof(f3143, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3142, f466])).
fof(f3142, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3141, f467])).
fof(f3141, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3140, f468])).
fof(f3140, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3139, f469])).
fof(f3139, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3138, f470])).
fof(f3138, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3137, f471])).
fof(f3137, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3136, f472])).
fof(f3136, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3135, f473])).
fof(f3135, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3134, f474])).
fof(f3134, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3133, f475])).
fof(f3133, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3132, f476])).
fof(f3132, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3131, f477])).
fof(f3131, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3130, f478])).
fof(f3130, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3129, f479])).
fof(f3129, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3128, f2473])).
fof(f3128, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3127, f2927])).
fof(f3127, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3126, f2931])).
fof(f3126, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_3 | ~ spl43_69)), inference(subsumption_resolution, [], [f3104, f2935])).
fof(f3104, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_3), inference(resolution, [], [f933, f904])).
fof(f904, plain, ! [X2, X0, X1] : (v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v3_waybel34(X2, X1, X0) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f457])).
fof(f933, plain, (~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | spl43_3), inference(avatar_component_clause, [], [f931])).
fof(f3125, plain, (~ spl43_1 | spl43_71), inference(avatar_contradiction_clause, [], [f3124])).
fof(f3124, plain, ($false | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3123, f466])).
fof(f3123, plain, (~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3122, f467])).
fof(f3122, plain, (~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3121, f468])).
fof(f3121, plain, (~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3120, f469])).
fof(f3120, plain, (~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3119, f470])).
fof(f3119, plain, (~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3118, f471])).
fof(f3118, plain, (~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3117, f472])).
fof(f3117, plain, (~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3116, f473])).
fof(f3116, plain, (~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3115, f474])).
fof(f3115, plain, (~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3114, f475])).
fof(f3114, plain, (~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3113, f476])).
fof(f3113, plain, (~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3112, f477])).
fof(f3112, plain, (~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3111, f478])).
fof(f3111, plain, (~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3110, f479])).
fof(f3110, plain, (~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3109, f481])).
fof(f3109, plain, (~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3108, f482])).
fof(f3108, plain, (~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3107, f483])).
fof(f483, plain, v17_waybel_0(sK9, sK7, sK8), inference(cnf_transformation, [], [f388])).
fof(f3107, plain, (~ v17_waybel_0(sK9, sK7, sK8) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3106, f484])).
fof(f3106, plain, (~ m2_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v17_waybel_0(sK9, sK7, sK8) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_1 | spl43_71)), inference(subsumption_resolution, [], [f3105, f924])).
fof(f924, plain, (v22_waybel_0(sK9, sK7, sK8) | ~ spl43_1), inference(avatar_component_clause, [], [f923])).
fof(f923, plain, (spl43_1 <=> v22_waybel_0(sK9, sK7, sK8)), introduced(avatar_definition, [new_symbols(naming, [spl43_1])])).
fof(f3105, plain, (~ v22_waybel_0(sK9, sK7, sK8) | ~ m2_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v17_waybel_0(sK9, sK7, sK8) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_71), inference(resolution, [], [f3083, f894])).
fof(f894, plain, ! [X2, X0, X1] : (v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ v22_waybel_0(X2, X0, X1) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f350])).
fof(f350, plain, ! [X0] : (! [X1] : (! [X2] : (v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ v22_waybel_0(X2, X0, X1) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f349])).
fof(f349, plain, ! [X0] : (! [X1] : (! [X2] : ((v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ v22_waybel_0(X2, X0, X1)) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f127])).
fof(f127, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => (v22_waybel_0(X2, X0, X1) => v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT380+1.p', t24_waybel34)).
fof(f3083, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | spl43_71), inference(avatar_component_clause, [], [f3081])).
fof(f3081, plain, (spl43_71 <=> v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7)), introduced(avatar_definition, [new_symbols(naming, [spl43_71])])).
fof(f3103, plain, (spl43_71 | ~ spl43_70 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69), inference(avatar_split_clause, [], [f3102, f2934, f2930, f2926, f2938, f3081])).
fof(f3102, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3101, f466])).
fof(f3101, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3100, f467])).
fof(f3100, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3099, f468])).
fof(f3099, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3098, f469])).
fof(f3098, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3097, f470])).
fof(f3097, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3096, f471])).
fof(f3096, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3095, f472])).
fof(f3095, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3094, f473])).
fof(f3094, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3093, f474])).
fof(f3093, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3092, f475])).
fof(f3092, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3091, f476])).
fof(f3091, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3090, f477])).
fof(f3090, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3089, f478])).
fof(f3089, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3088, f479])).
fof(f3088, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3087, f2473])).
fof(f3087, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3086, f2927])).
fof(f3086, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3085, f2931])).
fof(f3085, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | ~ spl43_69), inference(subsumption_resolution, [], [f2998, f480])).
fof(f480, plain, v2_waybel_8(sK8), inference(cnf_transformation, [], [f388])).
fof(f2998, plain, (~ v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_waybel_8(sK8) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | ~ spl43_69), inference(resolution, [], [f2935, f901])).
fof(f901, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v3_waybel34(X2, X1, X0) | ~ v2_waybel_8(X1) | v1_waybel34(X2, X1, X0) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f362])).
fof(f362, plain, ! [X0] : (! [X1] : (! [X2] : (v1_waybel34(X2, X1, X0) | ~ v3_waybel34(X2, X1, X0) | ~ v2_waybel_8(X1) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f361])).
fof(f361, plain, ! [X0] : (! [X1] : (! [X2] : ((v1_waybel34(X2, X1, X0) | (~ v3_waybel34(X2, X1, X0) | ~ v2_waybel_8(X1))) | (~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f134])).
fof(f134, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2)) => ((v3_waybel34(X2, X1, X0) & v2_waybel_8(X1)) => v1_waybel34(X2, X1, X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT380+1.p', t61_waybel34)).
fof(f3084, plain, (spl43_70 | ~ spl43_71 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69), inference(avatar_split_clause, [], [f3079, f2934, f2930, f2926, f3081, f2938])).
fof(f3079, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3078, f466])).
fof(f3078, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3077, f467])).
fof(f3077, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3076, f468])).
fof(f3076, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3075, f469])).
fof(f3075, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3074, f470])).
fof(f3074, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3073, f471])).
fof(f3073, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3072, f472])).
fof(f3072, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3071, f473])).
fof(f3071, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3070, f474])).
fof(f3070, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3069, f475])).
fof(f3069, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3068, f476])).
fof(f3068, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3067, f477])).
fof(f3067, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3066, f478])).
fof(f3066, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3065, f479])).
fof(f3065, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3064, f2473])).
fof(f3064, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3063, f2927])).
fof(f3063, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f2999, f2931])).
fof(f2999, plain, (~ v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | ~ spl43_69), inference(resolution, [], [f2935, f900])).
fof(f900, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_waybel34(X2, X1, X0) | v3_waybel34(X2, X1, X0) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f360])).
fof(f360, plain, ! [X0] : (! [X1] : (! [X2] : (v3_waybel34(X2, X1, X0) | ~ v1_waybel34(X2, X1, X0) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f359])).
fof(f359, plain, ! [X0] : (! [X1] : (! [X2] : ((v3_waybel34(X2, X1, X0) | ~ v1_waybel34(X2, X1, X0)) | (~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f133])).
fof(f133, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2)) => (v1_waybel34(X2, X1, X0) => v3_waybel34(X2, X1, X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT380+1.p', t60_waybel34)).
fof(f3062, plain, (~ spl43_3 | spl43_70 | ~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69), inference(avatar_split_clause, [], [f3061, f2934, f2930, f2926, f939, f935, f927, f2938, f931])).
fof(f3061, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3060, f466])).
fof(f3060, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3059, f467])).
fof(f3059, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3058, f468])).
fof(f3058, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3057, f469])).
fof(f3057, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3056, f470])).
fof(f3056, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3055, f471])).
fof(f3055, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3054, f472])).
fof(f3054, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3053, f473])).
fof(f3053, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3052, f474])).
fof(f3052, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3051, f475])).
fof(f3051, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3050, f476])).
fof(f3050, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3049, f477])).
fof(f3049, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3048, f478])).
fof(f3048, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3047, f479])).
fof(f3047, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3046, f2473])).
fof(f3046, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3045, f2927])).
fof(f3045, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_68 | ~ spl43_69)), inference(subsumption_resolution, [], [f3044, f2931])).
fof(f3044, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5 | ~ spl43_69)), inference(subsumption_resolution, [], [f3043, f2935])).
fof(f3043, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_2 | ~ spl43_4 | ~ spl43_5)), inference(subsumption_resolution, [], [f2907, f928])).
fof(f928, plain, (v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8)))) | ~ spl43_2), inference(avatar_component_clause, [], [f927])).
fof(f2907, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8)))) | ~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_4 | ~ spl43_5)), inference(subsumption_resolution, [], [f2904, f936])).
fof(f936, plain, (v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) | ~ spl43_4), inference(avatar_component_clause, [], [f935])).
fof(f2904, plain, (v3_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8)))) | ~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | ~ spl43_5), inference(resolution, [], [f907, f940])).
fof(f940, plain, (m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ spl43_5), inference(avatar_component_clause, [], [f939])).
fof(f907, plain, ! [X2, X0, X1] : (~ m2_relset_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | v3_waybel34(X2, X1, X0) | ~ v4_waybel34(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), k1_waybel_8(X1), k1_waybel_8(X0)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1))), u1_struct_0(k1_waybel_8(X1)), u1_struct_0(k1_waybel_8(X0))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(X1), u1_struct_0(X0), X2, u1_struct_0(k1_waybel_8(X1)))) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f457])).
fof(f3042, plain, (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70), inference(avatar_contradiction_clause, [], [f3041])).
fof(f3041, plain, ($false | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3040, f466])).
fof(f3040, plain, (~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3039, f467])).
fof(f3039, plain, (~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3038, f468])).
fof(f3038, plain, (~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3037, f469])).
fof(f3037, plain, (~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3036, f470])).
fof(f3036, plain, (~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3035, f471])).
fof(f3035, plain, (~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3034, f472])).
fof(f3034, plain, (~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3033, f473])).
fof(f3033, plain, (~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3032, f474])).
fof(f3032, plain, (~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3031, f475])).
fof(f3031, plain, (~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3030, f476])).
fof(f3030, plain, (~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3029, f477])).
fof(f3029, plain, (~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3028, f478])).
fof(f3028, plain, (~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3027, f1615])).
fof(f1615, plain, v3_waybel_3(sK8), inference(resolution, [], [f1605, f561])).
fof(f561, plain, ! [X0] : (~ sP1(X0) | v3_waybel_3(X0)), inference(cnf_transformation, [], [f390])).
fof(f390, plain, ! [X0] : ((v3_waybel_3(X0) & v2_waybel_3(X0) & v24_waybel_0(X0) & v2_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) | ~ sP1(X0)), inference(nnf_transformation, [], [f372])).
fof(f372, plain, ! [X0] : ((v3_waybel_3(X0) & v2_waybel_3(X0) & v24_waybel_0(X0) & v2_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) | ~ sP1(X0)), inference(usedef, [], [e372])).
fof(e372, plain, ! [X0] : (sP1(X0) <=> (v3_waybel_3(X0) & v2_waybel_3(X0) & v24_waybel_0(X0) & v2_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f1605, plain, sP1(sK8), inference(subsumption_resolution, [], [f1604, f479])).
fof(f1604, plain, (sP1(sK8) | ~ l1_orders_2(sK8)), inference(subsumption_resolution, [], [f1603, f473])).
fof(f1603, plain, (sP1(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK8)), inference(subsumption_resolution, [], [f1602, f474])).
fof(f1602, plain, (sP1(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK8)), inference(subsumption_resolution, [], [f1601, f475])).
fof(f1601, plain, (sP1(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK8)), inference(subsumption_resolution, [], [f1600, f476])).
fof(f1600, plain, (sP1(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK8)), inference(subsumption_resolution, [], [f1599, f477])).
fof(f1599, plain, (sP1(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK8)), inference(resolution, [], [f562, f480])).
fof(f562, plain, ! [X0] : (~ v2_waybel_8(X0) | sP1(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f373])).
fof(f373, plain, ! [X0] : (sP1(X0) | ~ v2_waybel_8(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | ~ l1_orders_2(X0)), inference(definition_folding, [], [f225, e372])).
fof(f225, plain, ! [X0] : ((v3_waybel_3(X0) & v2_waybel_3(X0) & v24_waybel_0(X0) & v2_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) | ~ v2_waybel_8(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f224])).
fof(f224, plain, ! [X0] : (((v3_waybel_3(X0) & v2_waybel_3(X0) & v24_waybel_0(X0) & v2_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) | (~ v2_waybel_8(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ! [X0] : (l1_orders_2(X0) => ((v2_waybel_8(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v3_waybel_3(X0) & v2_waybel_3(X0) & v24_waybel_0(X0) & v2_yellow_0(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT380+1.p', cc1_waybel_8)).
fof(f3027, plain, (~ v3_waybel_3(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3026, f479])).
fof(f3026, plain, (~ l1_orders_2(sK8) | ~ v3_waybel_3(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3025, f481])).
fof(f3025, plain, (~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_waybel_3(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3024, f482])).
fof(f3024, plain, (~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_waybel_3(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3023, f483])).
fof(f3023, plain, (~ v17_waybel_0(sK9, sK7, sK8) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_waybel_3(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3022, f484])).
fof(f3022, plain, (~ m2_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v17_waybel_0(sK9, sK7, sK8) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_waybel_3(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (spl43_1 | ~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3021, f925])).
fof(f925, plain, (~ v22_waybel_0(sK9, sK7, sK8) | spl43_1), inference(avatar_component_clause, [], [f923])).
fof(f3021, plain, (v22_waybel_0(sK9, sK7, sK8) | ~ m2_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v17_waybel_0(sK9, sK7, sK8) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_waybel_3(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(resolution, [], [f3020, f895])).
fof(f895, plain, ! [X2, X0, X1] : (~ v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | v22_waybel_0(X2, X0, X1) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_waybel_3(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f352])).
fof(f352, plain, ! [X0] : (! [X1] : (! [X2] : (v22_waybel_0(X2, X0, X1) | ~ v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_waybel_3(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f351])).
fof(f351, plain, ! [X0] : (! [X1] : (! [X2] : ((v22_waybel_0(X2, X0, X1) | ~ v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0)) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v3_waybel_3(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f128])).
fof(f128, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_waybel_3(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => (v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) => v22_waybel_0(X2, X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT380+1.p', t25_waybel34)).
fof(f3020, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3019, f466])).
fof(f3019, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3018, f467])).
fof(f3018, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3017, f468])).
fof(f3017, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3016, f469])).
fof(f3016, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3015, f470])).
fof(f3015, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3014, f471])).
fof(f3014, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3013, f472])).
fof(f3013, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3012, f473])).
fof(f3012, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3011, f474])).
fof(f3011, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3010, f475])).
fof(f3010, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3009, f476])).
fof(f3009, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3008, f477])).
fof(f3008, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3007, f478])).
fof(f3007, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3006, f479])).
fof(f3006, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3005, f2473])).
fof(f3005, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_67 | ~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3004, f2927])).
fof(f3004, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_68 | ~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3003, f2931])).
fof(f3003, plain, (v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f3002, f480])).
fof(f3002, plain, (~ v2_waybel_8(sK8) | v1_waybel34(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | ~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | ~ v1_funct_1(k1_waybel34(sK7, sK8, sK9)) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | (~ spl43_69 | ~ spl43_70)), inference(subsumption_resolution, [], [f2998, f2940])).
fof(f2997, plain, spl43_69, inference(avatar_contradiction_clause, [], [f2996])).
fof(f2996, plain, ($false | spl43_69), inference(subsumption_resolution, [], [f2995, f466])).
fof(f2995, plain, (~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2994, f467])).
fof(f2994, plain, (~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2993, f468])).
fof(f2993, plain, (~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2992, f469])).
fof(f2992, plain, (~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2991, f470])).
fof(f2991, plain, (~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2990, f472])).
fof(f2990, plain, (~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2989, f473])).
fof(f2989, plain, (~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2988, f474])).
fof(f2988, plain, (~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2987, f475])).
fof(f2987, plain, (~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2986, f476])).
fof(f2986, plain, (~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2985, f477])).
fof(f2985, plain, (~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2984, f479])).
fof(f2984, plain, (~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2983, f481])).
fof(f2983, plain, (~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2982, f482])).
fof(f2982, plain, (~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(subsumption_resolution, [], [f2981, f1185])).
fof(f2981, plain, (~ m1_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_69), inference(resolution, [], [f2936, f657])).
fof(f657, plain, ! [X2, X0, X1] : (m2_relset_1(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f277])).
fof(f2936, plain, (~ m2_relset_1(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | spl43_69), inference(avatar_component_clause, [], [f2934])).
fof(f2980, plain, spl43_68, inference(avatar_contradiction_clause, [], [f2979])).
fof(f2979, plain, ($false | spl43_68), inference(subsumption_resolution, [], [f2978, f466])).
fof(f2978, plain, (~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2977, f467])).
fof(f2977, plain, (~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2976, f468])).
fof(f2976, plain, (~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2975, f469])).
fof(f2975, plain, (~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2974, f470])).
fof(f2974, plain, (~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2973, f471])).
fof(f2973, plain, (~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2972, f472])).
fof(f2972, plain, (~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2971, f473])).
fof(f2971, plain, (~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2970, f474])).
fof(f2970, plain, (~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2969, f475])).
fof(f2969, plain, (~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2968, f476])).
fof(f2968, plain, (~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2967, f477])).
fof(f2967, plain, (~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2966, f478])).
fof(f2966, plain, (~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2965, f479])).
fof(f2965, plain, (~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2964, f481])).
fof(f2964, plain, (~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2963, f482])).
fof(f2963, plain, (~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2962, f483])).
fof(f2962, plain, (~ v17_waybel_0(sK9, sK7, sK8) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(subsumption_resolution, [], [f2961, f1185])).
fof(f2961, plain, (~ m1_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v17_waybel_0(sK9, sK7, sK8) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v3_lattice3(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v3_lattice3(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_68), inference(resolution, [], [f2932, f708])).
fof(f708, plain, ! [X2, X0, X1] : (v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f300])).
fof(f300, plain, ! [X0, X1, X2] : ((v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f299])).
fof(f299, plain, ! [X0, X1, X2] : ((v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f172])).
fof(f172, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f166])).
fof(f166, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k1_waybel34(X0, X1, X2), X1, X0) & v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f151])).
fof(f151, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k1_waybel34(X0, X1, X2), X1, X0) & v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_partfun1(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f85])).
fof(f85, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k1_waybel34(X0, X1, X2), X1, X0) & v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v20_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_partfun1(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT380+1.p', fc1_waybel34)).
fof(f2932, plain, (~ v18_waybel_0(k1_waybel34(sK7, sK8, sK9), sK8, sK7) | spl43_68), inference(avatar_component_clause, [], [f2930])).
fof(f2960, plain, spl43_67, inference(avatar_contradiction_clause, [], [f2959])).
fof(f2959, plain, ($false | spl43_67), inference(subsumption_resolution, [], [f2958, f466])).
fof(f2958, plain, (~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2957, f467])).
fof(f2957, plain, (~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2956, f468])).
fof(f2956, plain, (~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2955, f469])).
fof(f2955, plain, (~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2954, f470])).
fof(f2954, plain, (~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2953, f472])).
fof(f2953, plain, (~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2952, f473])).
fof(f2952, plain, (~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2951, f474])).
fof(f2951, plain, (~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2950, f475])).
fof(f2950, plain, (~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2949, f476])).
fof(f2949, plain, (~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2948, f477])).
fof(f2948, plain, (~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2947, f479])).
fof(f2947, plain, (~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2946, f481])).
fof(f2946, plain, (~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2945, f482])).
fof(f2945, plain, (~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(subsumption_resolution, [], [f2944, f1185])).
fof(f2944, plain, (~ m1_relset_1(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_2(sK9, u1_struct_0(sK7), u1_struct_0(sK8)) | ~ v1_funct_1(sK9) | ~ l1_orders_2(sK8) | ~ v2_lattice3(sK8) | ~ v1_lattice3(sK8) | ~ v4_orders_2(sK8) | ~ v3_orders_2(sK8) | ~ v2_orders_2(sK8) | ~ l1_orders_2(sK7) | ~ v2_lattice3(sK7) | ~ v1_lattice3(sK7) | ~ v4_orders_2(sK7) | ~ v3_orders_2(sK7) | ~ v2_orders_2(sK7) | spl43_67), inference(resolution, [], [f2928, f656])).
fof(f656, plain, ! [X2, X0, X1] : (v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f277])).
fof(f2928, plain, (~ v1_funct_2(k1_waybel34(sK7, sK8, sK9), u1_struct_0(sK8), u1_struct_0(sK7)) | spl43_67), inference(avatar_component_clause, [], [f2926])).
fof(f946, plain, (spl43_1 | spl43_2), inference(avatar_split_clause, [], [f485, f927, f923])).
fof(f485, plain, (v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8)))) | v22_waybel_0(sK9, sK7, sK8)), inference(cnf_transformation, [], [f388])).
fof(f945, plain, (spl43_1 | spl43_3), inference(avatar_split_clause, [], [f486, f931, f923])).
fof(f486, plain, (v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | v22_waybel_0(sK9, sK7, sK8)), inference(cnf_transformation, [], [f388])).
fof(f944, plain, (spl43_1 | spl43_4), inference(avatar_split_clause, [], [f487, f935, f923])).
fof(f487, plain, (v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) | v22_waybel_0(sK9, sK7, sK8)), inference(cnf_transformation, [], [f388])).
fof(f943, plain, (spl43_1 | spl43_5), inference(avatar_split_clause, [], [f488, f939, f923])).
fof(f488, plain, (m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | v22_waybel_0(sK9, sK7, sK8)), inference(cnf_transformation, [], [f388])).
fof(f942, plain, (~ spl43_1 | ~ spl43_2 | ~ spl43_3 | ~ spl43_4 | ~ spl43_5), inference(avatar_split_clause, [], [f489, f939, f935, f931, f927, f923])).
fof(f489, plain, (~ m2_relset_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v4_waybel34(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), k1_waybel_8(sK8), k1_waybel_8(sK7)) | ~ v1_funct_2(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8))), u1_struct_0(k1_waybel_8(sK8)), u1_struct_0(k1_waybel_8(sK7))) | ~ v1_funct_1(k8_relset_1(u1_struct_0(sK8), u1_struct_0(sK7), k1_waybel34(sK7, sK8, sK9), u1_struct_0(k1_waybel_8(sK8)))) | ~ v22_waybel_0(sK9, sK7, sK8)), inference(cnf_transformation, [], [f388])).