fof(f3642, plain, $false, inference(avatar_sat_refutation, [], [f505, f959, f3568, f3587, f3637])).
fof(f3637, plain, ~ spl33_1, inference(avatar_contradiction_clause, [], [f3636])).
fof(f3636, plain, ($false | ~ spl33_1), inference(subsumption_resolution, [], [f3635, f256])).
fof(f256, plain, m1_subset_1(sK4, u1_struct_0(sK3)), inference(cnf_transformation, [], [f192])).
fof(f192, plain, ((((~ l3_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | ~ v17_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | ~ v10_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | v3_struct_0(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5)))) & r3_lattices(sK3, sK4, sK5) & l3_lattices(sK3) & v17_lattices(sK3) & v10_lattices(sK3) & ~ v3_struct_0(sK3) & m1_subset_1(sK5, u1_struct_0(sK3))) & m1_subset_1(sK4, u1_struct_0(sK3))) & l3_lattices(sK3) & v10_lattices(sK3) & ~ v3_struct_0(sK3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3, sK4, sK5])], [f98, f191, f190, f189])).
fof(f189, plain, (? [X0] : (? [X1] : (? [X2] : ((~ l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v10_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))) & r3_lattices(X0, X1, X2) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0) & m1_subset_1(X2, u1_struct_0(X0))) & m1_subset_1(X1, u1_struct_0(X0))) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (? [X1] : (? [X2] : ((~ l3_lattices(k23_filter_2(sK3, k22_filter_2(sK3, X1, X2))) | ~ v17_lattices(k23_filter_2(sK3, k22_filter_2(sK3, X1, X2))) | ~ v10_lattices(k23_filter_2(sK3, k22_filter_2(sK3, X1, X2))) | v3_struct_0(k23_filter_2(sK3, k22_filter_2(sK3, X1, X2)))) & r3_lattices(sK3, X1, X2) & l3_lattices(sK3) & v17_lattices(sK3) & v10_lattices(sK3) & ~ v3_struct_0(sK3) & m1_subset_1(X2, u1_struct_0(sK3))) & m1_subset_1(X1, u1_struct_0(sK3))) & l3_lattices(sK3) & v10_lattices(sK3) & ~ v3_struct_0(sK3))), introduced(choice_axiom, [])).
fof(f190, plain, (? [X1] : (? [X2] : ((~ l3_lattices(k23_filter_2(sK3, k22_filter_2(sK3, X1, X2))) | ~ v17_lattices(k23_filter_2(sK3, k22_filter_2(sK3, X1, X2))) | ~ v10_lattices(k23_filter_2(sK3, k22_filter_2(sK3, X1, X2))) | v3_struct_0(k23_filter_2(sK3, k22_filter_2(sK3, X1, X2)))) & r3_lattices(sK3, X1, X2) & l3_lattices(sK3) & v17_lattices(sK3) & v10_lattices(sK3) & ~ v3_struct_0(sK3) & m1_subset_1(X2, u1_struct_0(sK3))) & m1_subset_1(X1, u1_struct_0(sK3))) => (? [X2] : ((~ l3_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, X2))) | ~ v17_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, X2))) | ~ v10_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, X2))) | v3_struct_0(k23_filter_2(sK3, k22_filter_2(sK3, sK4, X2)))) & r3_lattices(sK3, sK4, X2) & l3_lattices(sK3) & v17_lattices(sK3) & v10_lattices(sK3) & ~ v3_struct_0(sK3) & m1_subset_1(X2, u1_struct_0(sK3))) & m1_subset_1(sK4, u1_struct_0(sK3)))), introduced(choice_axiom, [])).
fof(f191, plain, (? [X2] : ((~ l3_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, X2))) | ~ v17_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, X2))) | ~ v10_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, X2))) | v3_struct_0(k23_filter_2(sK3, k22_filter_2(sK3, sK4, X2)))) & r3_lattices(sK3, sK4, X2) & l3_lattices(sK3) & v17_lattices(sK3) & v10_lattices(sK3) & ~ v3_struct_0(sK3) & m1_subset_1(X2, u1_struct_0(sK3))) => ((~ l3_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | ~ v17_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | ~ v10_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | v3_struct_0(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5)))) & r3_lattices(sK3, sK4, sK5) & l3_lattices(sK3) & v17_lattices(sK3) & v10_lattices(sK3) & ~ v3_struct_0(sK3) & m1_subset_1(sK5, u1_struct_0(sK3)))), introduced(choice_axiom, [])).
fof(f98, plain, ? [X0] : (? [X1] : (? [X2] : ((~ l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v10_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))) & r3_lattices(X0, X1, X2) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0) & m1_subset_1(X2, u1_struct_0(X0))) & m1_subset_1(X1, u1_struct_0(X0))) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f97])).
fof(f97, plain, ? [X0] : (? [X1] : (? [X2] : (((~ l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v10_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))) & (r3_lattices(X0, X1, X2) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0))) & m1_subset_1(X2, u1_struct_0(X0))) & m1_subset_1(X1, u1_struct_0(X0))) & (l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => ((r3_lattices(X0, X1, X2) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v10_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & ~ v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => ((r3_lattices(X0, X1, X2) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v10_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & ~ v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT337+1.p', t87_filter_2)).
fof(f3635, plain, (~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ spl33_1), inference(subsumption_resolution, [], [f3634, f257])).
fof(f257, plain, m1_subset_1(sK5, u1_struct_0(sK3)), inference(cnf_transformation, [], [f192])).
fof(f3634, plain, (~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ spl33_1), inference(subsumption_resolution, [], [f3633, f258])).
fof(f258, plain, ~ v3_struct_0(sK3), inference(cnf_transformation, [], [f192])).
fof(f3633, plain, (v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ spl33_1), inference(subsumption_resolution, [], [f3632, f259])).
fof(f259, plain, v10_lattices(sK3), inference(cnf_transformation, [], [f192])).
fof(f3632, plain, (~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ spl33_1), inference(subsumption_resolution, [], [f3631, f689])).
fof(f689, plain, v15_lattices(sK3), inference(subsumption_resolution, [], [f688, f261])).
fof(f261, plain, l3_lattices(sK3), inference(cnf_transformation, [], [f192])).
fof(f688, plain, (v15_lattices(sK3) | ~ l3_lattices(sK3)), inference(subsumption_resolution, [], [f686, f258])).
fof(f686, plain, (v15_lattices(sK3) | v3_struct_0(sK3) | ~ l3_lattices(sK3)), inference(resolution, [], [f288, f260])).
fof(f260, plain, v17_lattices(sK3), inference(cnf_transformation, [], [f192])).
fof(f288, plain, ! [X0] : (~ v17_lattices(X0) | v15_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ! [X0] : ((v16_lattices(X0) & v15_lattices(X0) & v14_lattices(X0) & v13_lattices(X0) & v11_lattices(X0) & ~ v3_struct_0(X0)) | ~ v17_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f114])).
fof(f114, plain, ! [X0] : (((v16_lattices(X0) & v15_lattices(X0) & v14_lattices(X0) & v13_lattices(X0) & v11_lattices(X0) & ~ v3_struct_0(X0)) | (~ v17_lattices(X0) | v3_struct_0(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ! [X0] : (l3_lattices(X0) => ((v17_lattices(X0) & ~ v3_struct_0(X0)) => (v16_lattices(X0) & v15_lattices(X0) & v14_lattices(X0) & v13_lattices(X0) & v11_lattices(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT337+1.p', cc5_lattices)).
fof(f3631, plain, (~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ spl33_1), inference(subsumption_resolution, [], [f3630, f695])).
fof(f695, plain, v16_lattices(sK3), inference(subsumption_resolution, [], [f694, f261])).
fof(f694, plain, (v16_lattices(sK3) | ~ l3_lattices(sK3)), inference(subsumption_resolution, [], [f692, f258])).
fof(f692, plain, (v16_lattices(sK3) | v3_struct_0(sK3) | ~ l3_lattices(sK3)), inference(resolution, [], [f289, f260])).
fof(f289, plain, ! [X0] : (~ v17_lattices(X0) | v16_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f115])).
fof(f3630, plain, (~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ spl33_1), inference(subsumption_resolution, [], [f3629, f261])).
fof(f3629, plain, (~ l3_lattices(sK3) | ~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ spl33_1), inference(subsumption_resolution, [], [f3628, f828])).
fof(f828, plain, v12_lattices(sK3), inference(resolution, [], [f797, f300])).
fof(f300, plain, ! [X0] : (~ sP0(X0) | v12_lattices(X0)), inference(cnf_transformation, [], [f193])).
fof(f193, plain, ! [X0] : ((v12_lattices(X0) & v10_lattices(X0) & v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)) | ~ sP0(X0)), inference(nnf_transformation, [], [f183])).
fof(f183, plain, ! [X0] : ((v12_lattices(X0) & v10_lattices(X0) & v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)) | ~ sP0(X0)), inference(usedef, [], [e183])).
fof(e183, plain, ! [X0] : (sP0(X0) <=> (v12_lattices(X0) & v10_lattices(X0) & v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f797, plain, sP0(sK3), inference(subsumption_resolution, [], [f796, f261])).
fof(f796, plain, (sP0(sK3) | ~ l3_lattices(sK3)), inference(subsumption_resolution, [], [f795, f258])).
fof(f795, plain, (sP0(sK3) | v3_struct_0(sK3) | ~ l3_lattices(sK3)), inference(subsumption_resolution, [], [f788, f259])).
fof(f788, plain, (sP0(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ l3_lattices(sK3)), inference(resolution, [], [f301, f677])).
fof(f677, plain, v11_lattices(sK3), inference(subsumption_resolution, [], [f676, f261])).
fof(f676, plain, (v11_lattices(sK3) | ~ l3_lattices(sK3)), inference(subsumption_resolution, [], [f674, f258])).
fof(f674, plain, (v11_lattices(sK3) | v3_struct_0(sK3) | ~ l3_lattices(sK3)), inference(resolution, [], [f285, f260])).
fof(f285, plain, ! [X0] : (~ v17_lattices(X0) | v11_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f115])).
fof(f301, plain, ! [X0] : (~ v11_lattices(X0) | sP0(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ! [X0] : (sP0(X0) | ~ v11_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(definition_folding, [], [f119, e183])).
fof(f119, plain, ! [X0] : ((v12_lattices(X0) & v10_lattices(X0) & v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)) | ~ v11_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f118])).
fof(f118, plain, ! [X0] : (((v12_lattices(X0) & v10_lattices(X0) & v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)) | (~ v11_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ! [X0] : (l3_lattices(X0) => ((v11_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v12_lattices(X0) & v10_lattices(X0) & v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT337+1.p', cc7_lattices)).
fof(f3628, plain, (~ v12_lattices(sK3) | ~ l3_lattices(sK3) | ~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ spl33_1), inference(subsumption_resolution, [], [f3625, f262])).
fof(f262, plain, r3_lattices(sK3, sK4, sK5), inference(cnf_transformation, [], [f192])).
fof(f3625, plain, (~ r3_lattices(sK3, sK4, sK5) | ~ v12_lattices(sK3) | ~ l3_lattices(sK3) | ~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ spl33_1), inference(resolution, [], [f492, f484])).
fof(f484, plain, ! [X2, X0, X1] : (~ v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0))), inference(duplicate_literal_removal, [], [f478])).
fof(f478, plain, ! [X2, X0, X1] : (~ v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ! [X0] : (! [X1] : (! [X2] : ((l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v16_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v15_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v10_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & ~ v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f180])).
fof(f180, plain, ! [X0] : (! [X1] : (! [X2] : (((l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v16_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v15_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v10_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & ~ v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))) | (~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => ((r3_lattices(X0, X1, X2) & v12_lattices(X0) & l3_lattices(X0) & v16_lattices(X0) & v15_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v16_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v15_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & v10_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) & ~ v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT337+1.p', t86_filter_2)).
fof(f492, plain, (v3_struct_0(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | ~ spl33_1), inference(avatar_component_clause, [], [f490])).
fof(f490, plain, (spl33_1 <=> v3_struct_0(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl33_1])])).
fof(f3587, plain, spl33_4, inference(avatar_contradiction_clause, [], [f3586])).
fof(f3586, plain, ($false | spl33_4), inference(subsumption_resolution, [], [f3585, f256])).
fof(f3585, plain, (~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_4), inference(subsumption_resolution, [], [f3584, f257])).
fof(f3584, plain, (~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_4), inference(subsumption_resolution, [], [f3583, f258])).
fof(f3583, plain, (v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_4), inference(subsumption_resolution, [], [f3582, f259])).
fof(f3582, plain, (~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_4), inference(subsumption_resolution, [], [f3581, f689])).
fof(f3581, plain, (~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_4), inference(subsumption_resolution, [], [f3580, f695])).
fof(f3580, plain, (~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_4), inference(subsumption_resolution, [], [f3579, f261])).
fof(f3579, plain, (~ l3_lattices(sK3) | ~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_4), inference(subsumption_resolution, [], [f3578, f828])).
fof(f3578, plain, (~ v12_lattices(sK3) | ~ l3_lattices(sK3) | ~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_4), inference(subsumption_resolution, [], [f3576, f262])).
fof(f3576, plain, (~ r3_lattices(sK3, sK4, sK5) | ~ v12_lattices(sK3) | ~ l3_lattices(sK3) | ~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_4), inference(resolution, [], [f504, f488])).
fof(f488, plain, ! [X2, X0, X1] : (l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0))), inference(duplicate_literal_removal, [], [f482])).
fof(f482, plain, ! [X2, X0, X1] : (l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f181])).
fof(f504, plain, (~ l3_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | spl33_4), inference(avatar_component_clause, [], [f502])).
fof(f502, plain, (spl33_4 <=> l3_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl33_4])])).
fof(f3568, plain, spl33_3, inference(avatar_contradiction_clause, [], [f3567])).
fof(f3567, plain, ($false | spl33_3), inference(subsumption_resolution, [], [f3566, f677])).
fof(f3566, plain, (~ v11_lattices(sK3) | spl33_3), inference(subsumption_resolution, [], [f3565, f262])).
fof(f3565, plain, (~ r3_lattices(sK3, sK4, sK5) | ~ v11_lattices(sK3) | spl33_3), inference(subsumption_resolution, [], [f3564, f828])).
fof(f3564, plain, (~ v12_lattices(sK3) | ~ r3_lattices(sK3, sK4, sK5) | ~ v11_lattices(sK3) | spl33_3), inference(subsumption_resolution, [], [f3563, f256])).
fof(f3563, plain, (~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ v12_lattices(sK3) | ~ r3_lattices(sK3, sK4, sK5) | ~ v11_lattices(sK3) | spl33_3), inference(subsumption_resolution, [], [f3562, f257])).
fof(f3562, plain, (~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ v12_lattices(sK3) | ~ r3_lattices(sK3, sK4, sK5) | ~ v11_lattices(sK3) | spl33_3), inference(subsumption_resolution, [], [f3561, f258])).
fof(f3561, plain, (v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ v12_lattices(sK3) | ~ r3_lattices(sK3, sK4, sK5) | ~ v11_lattices(sK3) | spl33_3), inference(subsumption_resolution, [], [f3560, f259])).
fof(f3560, plain, (~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ v12_lattices(sK3) | ~ r3_lattices(sK3, sK4, sK5) | ~ v11_lattices(sK3) | spl33_3), inference(subsumption_resolution, [], [f3559, f689])).
fof(f3559, plain, (~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ v12_lattices(sK3) | ~ r3_lattices(sK3, sK4, sK5) | ~ v11_lattices(sK3) | spl33_3), inference(subsumption_resolution, [], [f3558, f695])).
fof(f3558, plain, (~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ v12_lattices(sK3) | ~ r3_lattices(sK3, sK4, sK5) | ~ v11_lattices(sK3) | spl33_3), inference(subsumption_resolution, [], [f3552, f261])).
fof(f3552, plain, (~ l3_lattices(sK3) | ~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | ~ v12_lattices(sK3) | ~ r3_lattices(sK3, sK4, sK5) | ~ v11_lattices(sK3) | spl33_3), inference(resolution, [], [f1605, f500])).
fof(f500, plain, (~ v17_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | spl33_3), inference(avatar_component_clause, [], [f498])).
fof(f498, plain, (spl33_3 <=> v17_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl33_3])])).
fof(f1605, plain, ! [X2, X0, X1] : (v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X2, X1))) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ v12_lattices(X0) | ~ r3_lattices(X0, X2, X1) | ~ v11_lattices(X0)), inference(subsumption_resolution, [], [f1604, f305])).
fof(f305, plain, ! [X2, X0, X1] : (m2_lattice4(k22_filter_2(X0, X1, X2), X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ! [X0, X1, X2] : ((m2_lattice4(k22_filter_2(X0, X1, X2), X0) & ~ v1_xboole_0(k22_filter_2(X0, X1, X2))) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f122])).
fof(f122, plain, ! [X0, X1, X2] : ((m2_lattice4(k22_filter_2(X0, X1, X2), X0) & ~ v1_xboole_0(k22_filter_2(X0, X1, X2))) | (~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0, X1, X2] : ((m1_subset_1(X2, u1_struct_0(X0)) & m1_subset_1(X1, u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (m2_lattice4(k22_filter_2(X0, X1, X2), X0) & ~ v1_xboole_0(k22_filter_2(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT337+1.p', dt_k22_filter_2)).
fof(f1604, plain, ! [X2, X0, X1] : (~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X2, X1))) | ~ r3_lattices(X0, X2, X1) | ~ v11_lattices(X0) | ~ m2_lattice4(k22_filter_2(X0, X2, X1), X0)), inference(subsumption_resolution, [], [f1603, f304])).
fof(f304, plain, ! [X2, X0, X1] : (~ v1_xboole_0(k22_filter_2(X0, X1, X2)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f123])).
fof(f1603, plain, ! [X2, X0, X1] : (~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X2, X1))) | ~ r3_lattices(X0, X2, X1) | ~ v11_lattices(X0) | ~ m2_lattice4(k22_filter_2(X0, X2, X1), X0) | v1_xboole_0(k22_filter_2(X0, X2, X1))), inference(duplicate_literal_removal, [], [f1602])).
fof(f1602, plain, ! [X2, X0, X1] : (~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X2, X1))) | ~ r3_lattices(X0, X2, X1) | ~ v11_lattices(X0) | ~ m2_lattice4(k22_filter_2(X0, X2, X1), X0) | v1_xboole_0(k22_filter_2(X0, X2, X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f970, f476])).
fof(f476, plain, ! [X0, X1] : (v11_lattices(k23_filter_2(X0, X1)) | ~ v11_lattices(X0) | ~ m2_lattice4(X1, X0) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f178])).
fof(f178, plain, ! [X0] : (! [X1] : (v11_lattices(k23_filter_2(X0, X1)) | ~ v11_lattices(X0) | ~ m2_lattice4(X1, X0) | v1_xboole_0(X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f177])).
fof(f177, plain, ! [X0] : (! [X1] : ((v11_lattices(k23_filter_2(X0, X1)) | ~ v11_lattices(X0)) | (~ m2_lattice4(X1, X0) | v1_xboole_0(X1))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f81])).
fof(f81, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) => (v11_lattices(X0) => v11_lattices(k23_filter_2(X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT337+1.p', t78_filter_2)).
fof(f970, plain, ! [X2, X0, X1] : (~ v11_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2)), inference(subsumption_resolution, [], [f969, f488])).
fof(f969, plain, ! [X2, X0, X1] : (~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v11_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))), inference(subsumption_resolution, [], [f968, f484])).
fof(f968, plain, ! [X2, X0, X1] : (~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v11_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))), inference(subsumption_resolution, [], [f967, f486])).
fof(f486, plain, ! [X2, X0, X1] : (v15_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0))), inference(duplicate_literal_removal, [], [f480])).
fof(f480, plain, ! [X2, X0, X1] : (v15_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f181])).
fof(f967, plain, ! [X2, X0, X1] : (~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | v17_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v15_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ v11_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | v3_struct_0(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ l3_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2)))), inference(resolution, [], [f487, f291])).
fof(f291, plain, ! [X0] : (~ v16_lattices(X0) | v17_lattices(X0) | ~ v15_lattices(X0) | ~ v11_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ! [X0] : ((v17_lattices(X0) & ~ v3_struct_0(X0)) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v11_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f116])).
fof(f116, plain, ! [X0] : (((v17_lattices(X0) & ~ v3_struct_0(X0)) | (~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v11_lattices(X0) | v3_struct_0(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0] : (l3_lattices(X0) => ((v16_lattices(X0) & v15_lattices(X0) & v11_lattices(X0) & ~ v3_struct_0(X0)) => (v17_lattices(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT337+1.p', cc6_lattices)).
fof(f487, plain, ! [X2, X0, X1] : (v16_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0))), inference(duplicate_literal_removal, [], [f481])).
fof(f481, plain, ! [X2, X0, X1] : (v16_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f181])).
fof(f959, plain, spl33_2, inference(avatar_contradiction_clause, [], [f958])).
fof(f958, plain, ($false | spl33_2), inference(subsumption_resolution, [], [f957, f256])).
fof(f957, plain, (~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_2), inference(subsumption_resolution, [], [f956, f257])).
fof(f956, plain, (~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_2), inference(subsumption_resolution, [], [f955, f258])).
fof(f955, plain, (v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_2), inference(subsumption_resolution, [], [f954, f259])).
fof(f954, plain, (~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_2), inference(subsumption_resolution, [], [f953, f689])).
fof(f953, plain, (~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_2), inference(subsumption_resolution, [], [f952, f695])).
fof(f952, plain, (~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_2), inference(subsumption_resolution, [], [f951, f261])).
fof(f951, plain, (~ l3_lattices(sK3) | ~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_2), inference(subsumption_resolution, [], [f950, f828])).
fof(f950, plain, (~ v12_lattices(sK3) | ~ l3_lattices(sK3) | ~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_2), inference(subsumption_resolution, [], [f949, f262])).
fof(f949, plain, (~ r3_lattices(sK3, sK4, sK5) | ~ v12_lattices(sK3) | ~ l3_lattices(sK3) | ~ v16_lattices(sK3) | ~ v15_lattices(sK3) | ~ v10_lattices(sK3) | v3_struct_0(sK3) | ~ m1_subset_1(sK5, u1_struct_0(sK3)) | ~ m1_subset_1(sK4, u1_struct_0(sK3)) | spl33_2), inference(resolution, [], [f485, f496])).
fof(f496, plain, (~ v10_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | spl33_2), inference(avatar_component_clause, [], [f494])).
fof(f494, plain, (spl33_2 <=> v10_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl33_2])])).
fof(f485, plain, ! [X2, X0, X1] : (v10_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0))), inference(duplicate_literal_removal, [], [f479])).
fof(f479, plain, ! [X2, X0, X1] : (v10_lattices(k23_filter_2(X0, k22_filter_2(X0, X1, X2))) | ~ r3_lattices(X0, X1, X2) | ~ v12_lattices(X0) | ~ l3_lattices(X0) | ~ v16_lattices(X0) | ~ v15_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f181])).
fof(f505, plain, (spl33_1 | ~ spl33_2 | ~ spl33_3 | ~ spl33_4), inference(avatar_split_clause, [], [f263, f502, f498, f494, f490])).
fof(f263, plain, (~ l3_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | ~ v17_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | ~ v10_lattices(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5))) | v3_struct_0(k23_filter_2(sK3, k22_filter_2(sK3, sK4, sK5)))), inference(cnf_transformation, [], [f192])).