fof(f45304, plain, $false, inference(avatar_sat_refutation, [], [f277, f282, f292, f297, f1589, f4448, f4800, f11848, f26850, f26879, f27644, f27678, f45282])).
fof(f45282, plain, (spl23_6 | ~ spl23_7 | ~ spl23_80 | ~ spl23_744), inference(avatar_contradiction_clause, [], [f45281])).
fof(f45281, plain, ($false | (spl23_6 | ~ spl23_7 | ~ spl23_80 | ~ spl23_744)), inference(subsumption_resolution, [], [f45280, f1439])).
fof(f1439, plain, (sdtlseqdt0(sK22, xp) | ~ spl23_80), inference(avatar_component_clause, [], [f1437])).
fof(f1437, plain, (spl23_80 <=> sdtlseqdt0(sK22, xp)), introduced(avatar_definition, [new_symbols(naming, [spl23_80])])).
fof(f45280, plain, (~ sdtlseqdt0(sK22, xp) | (spl23_6 | ~ spl23_7 | ~ spl23_744)), inference(subsumption_resolution, [], [f45255, f240])).
fof(f240, plain, aElementOf0(xp, xU), inference(cnf_transformation, [], [f137])).
fof(f137, plain, (aInfimumOfIn0(xp, xP, xU) & ! [X0] : (sdtlseqdt0(X0, xp) | (~ aLowerBoundOfIn0(X0, xP, xU) & ((~ sdtlseqdt0(X0, sK20(X0)) & aElementOf0(sK20(X0), xP)) | ~ aElementOf0(X0, xU)))) & aLowerBoundOfIn0(xp, xP, xU) & ! [X2] : (sdtlseqdt0(xp, X2) | ~ aElementOf0(X2, xP)) & aElementOf0(xp, xU) & aElementOf0(xp, xU)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK20])], [f71, f136])).
fof(f136, plain, ! [X0] : (? [X1] : (~ sdtlseqdt0(X0, X1) & aElementOf0(X1, xP)) => (~ sdtlseqdt0(X0, sK20(X0)) & aElementOf0(sK20(X0), xP))), introduced(choice_axiom, [])).
fof(f71, plain, (aInfimumOfIn0(xp, xP, xU) & ! [X0] : (sdtlseqdt0(X0, xp) | (~ aLowerBoundOfIn0(X0, xP, xU) & (? [X1] : (~ sdtlseqdt0(X0, X1) & aElementOf0(X1, xP)) | ~ aElementOf0(X0, xU)))) & aLowerBoundOfIn0(xp, xP, xU) & ! [X2] : (sdtlseqdt0(xp, X2) | ~ aElementOf0(X2, xP)) & aElementOf0(xp, xU) & aElementOf0(xp, xU)), inference(ennf_transformation, [], [f37])).
fof(f37, plain, (aInfimumOfIn0(xp, xP, xU) & ! [X0] : ((aLowerBoundOfIn0(X0, xP, xU) | (! [X1] : (aElementOf0(X1, xP) => sdtlseqdt0(X0, X1)) & aElementOf0(X0, xU))) => sdtlseqdt0(X0, xp)) & aLowerBoundOfIn0(xp, xP, xU) & ! [X2] : (aElementOf0(X2, xP) => sdtlseqdt0(xp, X2)) & aElementOf0(xp, xU) & aElementOf0(xp, xU)), inference(rectify, [], [f28])).
fof(f28, plain, (aInfimumOfIn0(xp, xP, xU) & ! [X0] : ((aLowerBoundOfIn0(X0, xP, xU) | (! [X1] : (aElementOf0(X1, xP) => sdtlseqdt0(X0, X1)) & aElementOf0(X0, xU))) => sdtlseqdt0(X0, xp)) & aLowerBoundOfIn0(xp, xP, xU) & ! [X0] : (aElementOf0(X0, xP) => sdtlseqdt0(xp, X0)) & aElementOf0(xp, xU) & aElementOf0(xp, xU)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT386+4.p', m__1261)).
fof(f45255, plain, (~ aElementOf0(xp, xU) | ~ sdtlseqdt0(sK22, xp) | (spl23_6 | ~ spl23_7 | ~ spl23_744)), inference(resolution, [], [f28323, f291])).
fof(f291, plain, (~ sdtlseqdt0(sK22, sdtlpdtrp0(xf, xp)) | spl23_6), inference(avatar_component_clause, [], [f289])).
fof(f289, plain, (spl23_6 <=> sdtlseqdt0(sK22, sdtlpdtrp0(xf, xp))), introduced(avatar_definition, [new_symbols(naming, [spl23_6])])).
fof(f28323, plain, (! [X0] : (sdtlseqdt0(sK22, sdtlpdtrp0(xf, X0)) | ~ aElementOf0(X0, xU) | ~ sdtlseqdt0(sK22, X0)) | (~ spl23_7 | ~ spl23_744)), inference(subsumption_resolution, [], [f28271, f25900])).
fof(f25900, plain, (aElementOf0(sK22, xU) | ~ spl23_744), inference(avatar_component_clause, [], [f25899])).
fof(f25899, plain, (spl23_744 <=> aElementOf0(sK22, xU)), introduced(avatar_definition, [new_symbols(naming, [spl23_744])])).
fof(f28271, plain, (! [X0] : (sdtlseqdt0(sK22, sdtlpdtrp0(xf, X0)) | ~ aElementOf0(X0, xU) | ~ aElementOf0(sK22, xU) | ~ sdtlseqdt0(sK22, X0)) | ~ spl23_7), inference(superposition, [], [f261, f28091])).
fof(f28091, plain, ((sK22 = sdtlpdtrp0(xf, sK22)) | ~ spl23_7), inference(resolution, [], [f27653, f222])).
fof(f222, plain, ! [X0] : (~ aElementOf0(X0, xS) | (sdtlpdtrp0(xf, X0) = X0)), inference(cnf_transformation, [], [f67])).
fof(f67, plain, ((xS = cS1142(xf)) & ! [X0] : ((aElementOf0(X0, xS) | (~ aFixedPointOf0(X0, xf) & (~ (sdtlpdtrp0(xf, X0) = X0) | ~ aElementOf0(X0, szDzozmdt0(xf))))) & ((aFixedPointOf0(X0, xf) & (sdtlpdtrp0(xf, X0) = X0) & aElementOf0(X0, szDzozmdt0(xf))) | ~ aElementOf0(X0, xS))) & aSet0(xS)), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ((xS = cS1142(xf)) & ! [X0] : (((aFixedPointOf0(X0, xf) | ((sdtlpdtrp0(xf, X0) = X0) & aElementOf0(X0, szDzozmdt0(xf)))) => aElementOf0(X0, xS)) & (aElementOf0(X0, xS) => (aFixedPointOf0(X0, xf) & (sdtlpdtrp0(xf, X0) = X0) & aElementOf0(X0, szDzozmdt0(xf))))) & aSet0(xS)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT386+4.p', m__1144)).
fof(f27653, plain, (aElementOf0(sK22, xS) | ~ spl23_7), inference(resolution, [], [f296, f228])).
fof(f228, plain, ! [X0] : (~ aElementOf0(X0, xT) | aElementOf0(X0, xS)), inference(cnf_transformation, [], [f68])).
fof(f68, plain, (aSubsetOf0(xT, xS) & ! [X0] : (aElementOf0(X0, xS) | ~ aElementOf0(X0, xT)) & aSet0(xT)), inference(ennf_transformation, [], [f26])).
fof(f26, plain, (aSubsetOf0(xT, xS) & ! [X0] : (aElementOf0(X0, xT) => aElementOf0(X0, xS)) & aSet0(xT)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT386+4.p', m__1173)).
fof(f296, plain, (aElementOf0(sK22, xT) | ~ spl23_7), inference(avatar_component_clause, [], [f294])).
fof(f294, plain, (spl23_7 <=> aElementOf0(sK22, xT)), introduced(avatar_definition, [new_symbols(naming, [spl23_7])])).
fof(f261, plain, ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ aElementOf0(X1, xU) | ~ aElementOf0(X0, xU) | ~ sdtlseqdt0(X0, X1)), inference(forward_demodulation, [], [f260, f259])).
fof(f259, plain, (xU = szDzozmdt0(xf)), inference(forward_demodulation, [], [f217, f218])).
fof(f218, plain, (xU = szRzazndt0(xf)), inference(cnf_transformation, [], [f133])).
fof(f133, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (sP2(X2) | (~ aSubsetOf0(X2, xU) & ((~ aElementOf0(sK18(X2), xU) & aElementOf0(sK18(X2), X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK18])], [f76, f132])).
fof(f132, plain, ! [X2] : (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) => (~ aElementOf0(sK18(X2), xU) & aElementOf0(sK18(X2), X2))), introduced(choice_axiom, [])).
fof(f76, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (sP2(X2) | (~ aSubsetOf0(X2, xU) & (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(definition_folding, [], [f66, e75, e74, e73])).
fof(f73, plain, ! [X2] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) | ~ sP0(X2)), inference(usedef, [], [e73])).
fof(e73, plain, ! [X2] : (sP0(X2) <=> ? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f74, plain, ! [X4, X2] : (! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU)))) | ~ sP1(X4, X2)), inference(usedef, [], [e74])).
fof(e74, plain, ! [X4, X2] : (sP1(X4, X2) <=> ! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f75, plain, ! [X2] : (? [X4] : (sP0(X2) & aInfimumOfIn0(X4, X2, xU) & sP1(X4, X2) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU)) | ~ sP2(X2)), inference(usedef, [], [e75])).
fof(e75, plain, ! [X2] : (sP2(X2) <=> ? [X4] : (sP0(X2) & aInfimumOfIn0(X4, X2, xU) & sP1(X4, X2) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f66, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (? [X4] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) & aInfimumOfIn0(X4, X2, xU) & ! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU)))) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU)) | (~ aSubsetOf0(X2, xU) & (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(flattening, [], [f65])).
fof(f65, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : ((sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1)) | (~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (? [X4] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) & aInfimumOfIn0(X4, X2, xU) & ! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU)))) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU)) | (~ aSubsetOf0(X2, xU) & (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(ennf_transformation, [], [f35])).
fof(f35, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : ((aElementOf0(X1, szDzozmdt0(xf)) & aElementOf0(X0, szDzozmdt0(xf))) => (sdtlseqdt0(X0, X1) => sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : ((aSubsetOf0(X2, xU) | (! [X3] : (aElementOf0(X3, X2) => aElementOf0(X3, xU)) & aSet0(X2))) => ? [X4] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : ((aUpperBoundOfIn0(X6, X2, xU) | (! [X7] : (aElementOf0(X7, X2) => sdtlseqdt0(X7, X6)) & aElementOf0(X6, xU))) => sdtlseqdt0(X5, X6)) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (aElementOf0(X8, X2) => sdtlseqdt0(X8, X5)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) & aInfimumOfIn0(X4, X2, xU) & ! [X9] : ((aLowerBoundOfIn0(X9, X2, xU) | (! [X10] : (aElementOf0(X10, X2) => sdtlseqdt0(X9, X10)) & aElementOf0(X9, xU))) => sdtlseqdt0(X9, X4)) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (aElementOf0(X11, X2) => sdtlseqdt0(X4, X11)) & aElementOf0(X4, xU) & aElementOf0(X4, xU))) & aSet0(xU)), inference(rectify, [], [f24])).
fof(f24, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : ((aElementOf0(X1, szDzozmdt0(xf)) & aElementOf0(X0, szDzozmdt0(xf))) => (sdtlseqdt0(X0, X1) => sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X0] : ((aSubsetOf0(X0, xU) | (! [X1] : (aElementOf0(X1, X0) => aElementOf0(X1, xU)) & aSet0(X0))) => ? [X1] : (? [X2] : (aSupremumOfIn0(X2, X0, xU) & ! [X3] : ((aUpperBoundOfIn0(X3, X0, xU) | (! [X4] : (aElementOf0(X4, X0) => sdtlseqdt0(X4, X3)) & aElementOf0(X3, xU))) => sdtlseqdt0(X2, X3)) & aUpperBoundOfIn0(X2, X0, xU) & ! [X3] : (aElementOf0(X3, X0) => sdtlseqdt0(X3, X2)) & aElementOf0(X2, xU) & aElementOf0(X2, xU)) & aInfimumOfIn0(X1, X0, xU) & ! [X2] : ((aLowerBoundOfIn0(X2, X0, xU) | (! [X3] : (aElementOf0(X3, X0) => sdtlseqdt0(X2, X3)) & aElementOf0(X2, xU))) => sdtlseqdt0(X2, X1)) & aLowerBoundOfIn0(X1, X0, xU) & ! [X2] : (aElementOf0(X2, X0) => sdtlseqdt0(X1, X2)) & aElementOf0(X1, xU) & aElementOf0(X1, xU))) & aSet0(xU)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT386+4.p', m__1123)).
fof(f217, plain, (szDzozmdt0(xf) = szRzazndt0(xf)), inference(cnf_transformation, [], [f133])).
fof(f260, plain, ! [X0, X1] : (~ aElementOf0(X1, xU) | sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X0, szDzozmdt0(xf))), inference(forward_demodulation, [], [f215, f259])).
fof(f215, plain, ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))), inference(cnf_transformation, [], [f133])).
fof(f27678, plain, (spl23_744 | ~ spl23_7 | ~ spl23_39), inference(avatar_split_clause, [], [f27659, f913, f294, f25899])).
fof(f913, plain, (spl23_39 <=> aSubsetOf0(xT, xU)), introduced(avatar_definition, [new_symbols(naming, [spl23_39])])).
fof(f27659, plain, (aElementOf0(sK22, xU) | (~ spl23_7 | ~ spl23_39)), inference(resolution, [], [f296, f5053])).
fof(f5053, plain, (! [X4] : (~ aElementOf0(X4, xT) | aElementOf0(X4, xU)) | ~ spl23_39), inference(subsumption_resolution, [], [f5049, f209])).
fof(f209, plain, aSet0(xU), inference(cnf_transformation, [], [f133])).
fof(f5049, plain, (! [X4] : (~ aElementOf0(X4, xT) | aElementOf0(X4, xU) | ~ aSet0(xU)) | ~ spl23_39), inference(resolution, [], [f914, f148])).
fof(f148, plain, ! [X0, X3, X1] : (~ aSubsetOf0(X1, X0) | ~ aElementOf0(X3, X1) | aElementOf0(X3, X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | (~ aElementOf0(sK5(X0, X1), X0) & aElementOf0(sK5(X0, X1), X1)) | ~ aSet0(X1)) & ((! [X3] : (aElementOf0(X3, X0) | ~ aElementOf0(X3, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5])], [f85, f86])).
fof(f86, plain, ! [X1, X0] : (? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) => (~ aElementOf0(sK5(X0, X1), X0) & aElementOf0(sK5(X0, X1), X1))), introduced(choice_axiom, [])).
fof(f85, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | ? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1)) & ((! [X3] : (aElementOf0(X3, X0) | ~ aElementOf0(X3, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(rectify, [], [f84])).
fof(f84, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | ? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1)) & ((! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(flattening, [], [f83])).
fof(f83, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | (? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1))) & ((! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(nnf_transformation, [], [f42])).
fof(f42, plain, ! [X0] : (! [X1] : (aSubsetOf0(X1, X0) <=> (! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1))) | ~ aSet0(X0)), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aSubsetOf0(X1, X0) <=> (! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, X0)) & aSet0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT386+4.p', mDefSub)).
fof(f914, plain, (aSubsetOf0(xT, xU) | ~ spl23_39), inference(avatar_component_clause, [], [f913])).
fof(f27644, plain, (spl23_3 | ~ spl23_4 | ~ spl23_410 | ~ spl23_412 | ~ spl23_414), inference(avatar_contradiction_clause, [], [f27643])).
fof(f27643, plain, ($false | (spl23_3 | ~ spl23_4 | ~ spl23_410 | ~ spl23_412 | ~ spl23_414)), inference(subsumption_resolution, [], [f27641, f11957])).
fof(f11957, plain, (aElementOf0(sK21, xU) | ~ spl23_412), inference(avatar_component_clause, [], [f11956])).
fof(f11956, plain, (spl23_412 <=> aElementOf0(sK21, xU)), introduced(avatar_definition, [new_symbols(naming, [spl23_412])])).
fof(f27641, plain, (~ aElementOf0(sK21, xU) | (spl23_3 | ~ spl23_4 | ~ spl23_410 | ~ spl23_414)), inference(resolution, [], [f27172, f1357])).
fof(f1357, plain, ! [X2] : (aElement0(sdtlpdtrp0(xf, X2)) | ~ aElementOf0(X2, xU)), inference(subsumption_resolution, [], [f1355, f209])).
fof(f1355, plain, ! [X2] : (~ aElementOf0(X2, xU) | aElement0(sdtlpdtrp0(xf, X2)) | ~ aSet0(xU)), inference(resolution, [], [f573, f144])).
fof(f144, plain, ! [X0, X1] : (~ aElementOf0(X1, X0) | aElement0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f40])).
fof(f40, plain, ! [X0] : (! [X1] : (aElement0(X1) | ~ aElementOf0(X1, X0)) | ~ aSet0(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aElementOf0(X1, X0) => aElement0(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT386+4.p', mEOfElem)).
fof(f573, plain, ! [X0] : (aElementOf0(sdtlpdtrp0(xf, X0), xU) | ~ aElementOf0(X0, xU)), inference(forward_demodulation, [], [f572, f259])).
fof(f572, plain, ! [X0] : (aElementOf0(sdtlpdtrp0(xf, X0), xU) | ~ aElementOf0(X0, szDzozmdt0(xf))), inference(subsumption_resolution, [], [f567, f214])).
fof(f214, plain, aFunction0(xf), inference(cnf_transformation, [], [f133])).
fof(f567, plain, ! [X0] : (aElementOf0(sdtlpdtrp0(xf, X0), xU) | ~ aElementOf0(X0, szDzozmdt0(xf)) | ~ aFunction0(xf)), inference(superposition, [], [f182, f218])).
fof(f182, plain, ! [X0, X1] : (aElementOf0(sdtlpdtrp0(X0, X1), szRzazndt0(X0)) | ~ aElementOf0(X1, szDzozmdt0(X0)) | ~ aFunction0(X0)), inference(cnf_transformation, [], [f61])).
fof(f61, plain, ! [X0] : (! [X1] : (aElementOf0(sdtlpdtrp0(X0, X1), szRzazndt0(X0)) | ~ aElementOf0(X1, szDzozmdt0(X0))) | ~ aFunction0(X0)), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ! [X0] : (aFunction0(X0) => ! [X1] : (aElementOf0(X1, szDzozmdt0(X0)) => aElementOf0(sdtlpdtrp0(X0, X1), szRzazndt0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT386+4.p', mImgSort)).
fof(f27172, plain, (~ aElement0(sdtlpdtrp0(xf, sK21)) | (spl23_3 | ~ spl23_4 | ~ spl23_410 | ~ spl23_414)), inference(subsumption_resolution, [], [f27171, f26831])).
fof(f26831, plain, (sdtlseqdt0(xp, sK21) | ~ spl23_4), inference(resolution, [], [f281, f241])).
fof(f241, plain, ! [X2] : (~ aElementOf0(X2, xP) | sdtlseqdt0(xp, X2)), inference(cnf_transformation, [], [f137])).
fof(f281, plain, (aElementOf0(sK21, xP) | ~ spl23_4), inference(avatar_component_clause, [], [f279])).
fof(f279, plain, (spl23_4 <=> aElementOf0(sK21, xP)), introduced(avatar_definition, [new_symbols(naming, [spl23_4])])).
fof(f27171, plain, (~ aElement0(sdtlpdtrp0(xf, sK21)) | ~ sdtlseqdt0(xp, sK21) | (spl23_3 | ~ spl23_4 | ~ spl23_410 | ~ spl23_414)), inference(subsumption_resolution, [], [f27170, f240])).
fof(f27170, plain, (~ aElement0(sdtlpdtrp0(xf, sK21)) | ~ aElementOf0(xp, xU) | ~ sdtlseqdt0(xp, sK21) | (spl23_3 | ~ spl23_4 | ~ spl23_410 | ~ spl23_414)), inference(subsumption_resolution, [], [f27169, f281])).
fof(f27169, plain, (~ aElement0(sdtlpdtrp0(xf, sK21)) | ~ aElementOf0(sK21, xP) | ~ aElementOf0(xp, xU) | ~ sdtlseqdt0(xp, sK21) | (spl23_3 | ~ spl23_410 | ~ spl23_414)), inference(subsumption_resolution, [], [f27168, f11844])).
fof(f11844, plain, (aElement0(sdtlpdtrp0(xf, xp)) | ~ spl23_410), inference(avatar_component_clause, [], [f11843])).
fof(f11843, plain, (spl23_410 <=> aElement0(sdtlpdtrp0(xf, xp))), introduced(avatar_definition, [new_symbols(naming, [spl23_410])])).
fof(f27168, plain, (~ aElement0(sdtlpdtrp0(xf, sK21)) | ~ aElement0(sdtlpdtrp0(xf, xp)) | ~ aElementOf0(sK21, xP) | ~ aElementOf0(xp, xU) | ~ sdtlseqdt0(xp, sK21) | (spl23_3 | ~ spl23_414)), inference(subsumption_resolution, [], [f27167, f11991])).
fof(f11991, plain, (aElement0(sK21) | ~ spl23_414), inference(avatar_component_clause, [], [f11990])).
fof(f11990, plain, (spl23_414 <=> aElement0(sK21)), introduced(avatar_definition, [new_symbols(naming, [spl23_414])])).
fof(f27167, plain, (~ aElement0(sK21) | ~ aElement0(sdtlpdtrp0(xf, sK21)) | ~ aElement0(sdtlpdtrp0(xf, xp)) | ~ aElementOf0(sK21, xP) | ~ aElementOf0(xp, xU) | ~ sdtlseqdt0(xp, sK21) | spl23_3), inference(resolution, [], [f276, f3093])).
fof(f3093, plain, ! [X2, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X1), X2) | ~ aElement0(X2) | ~ aElement0(sdtlpdtrp0(xf, X2)) | ~ aElement0(sdtlpdtrp0(xf, X1)) | ~ aElementOf0(X2, xP) | ~ aElementOf0(X1, xU) | ~ sdtlseqdt0(X1, X2)), inference(subsumption_resolution, [], [f3075, f231])).
fof(f231, plain, ! [X0] : (~ aElementOf0(X0, xP) | aElementOf0(X0, xU)), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((aElementOf0(X0, xP) | (~ aUpperBoundOfIn0(X0, xT, xU) & (~ sdtlseqdt0(sK19(X0), X0) & aElementOf0(sK19(X0), xT))) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xU)) & ((aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (sdtlseqdt0(X2, X0) | ~ aElementOf0(X2, xT)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) | ~ aElementOf0(X0, xP))) & aSet0(xP)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK19])], [f70, f134])).
fof(f134, plain, ! [X0] : (? [X1] : (~ sdtlseqdt0(X1, X0) & aElementOf0(X1, xT)) => (~ sdtlseqdt0(sK19(X0), X0) & aElementOf0(sK19(X0), xT))), introduced(choice_axiom, [])).
fof(f70, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((aElementOf0(X0, xP) | (~ aUpperBoundOfIn0(X0, xT, xU) & ? [X1] : (~ sdtlseqdt0(X1, X0) & aElementOf0(X1, xT))) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xU)) & ((aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (sdtlseqdt0(X2, X0) | ~ aElementOf0(X2, xT)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) | ~ aElementOf0(X0, xP))) & aSet0(xP)), inference(flattening, [], [f69])).
fof(f69, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((aElementOf0(X0, xP) | ((~ aUpperBoundOfIn0(X0, xT, xU) & ? [X1] : (~ sdtlseqdt0(X1, X0) & aElementOf0(X1, xT))) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xU))) & ((aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (sdtlseqdt0(X2, X0) | ~ aElementOf0(X2, xT)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) | ~ aElementOf0(X0, xP))) & aSet0(xP)), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((((aUpperBoundOfIn0(X0, xT, xU) | ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0))) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) => aElementOf0(X0, xP)) & (aElementOf0(X0, xP) => (aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (aElementOf0(X2, xT) => sdtlseqdt0(X2, X0)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)))) & aSet0(xP)), inference(rectify, [], [f27])).
fof(f27, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((((aUpperBoundOfIn0(X0, xT, xU) | ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0))) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) => aElementOf0(X0, xP)) & (aElementOf0(X0, xP) => (aUpperBoundOfIn0(X0, xT, xU) & ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)))) & aSet0(xP)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT386+4.p', m__1244)).
fof(f3075, plain, ! [X2, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X1), X2) | ~ aElement0(X2) | ~ aElement0(sdtlpdtrp0(xf, X2)) | ~ aElement0(sdtlpdtrp0(xf, X1)) | ~ aElementOf0(X2, xP) | ~ aElementOf0(X2, xU) | ~ aElementOf0(X1, xU) | ~ sdtlseqdt0(X1, X2)), inference(resolution, [], [f786, f261])).
fof(f786, plain, ! [X6, X5] : (~ sdtlseqdt0(X5, sdtlpdtrp0(xf, X6)) | sdtlseqdt0(X5, X6) | ~ aElement0(X6) | ~ aElement0(sdtlpdtrp0(xf, X6)) | ~ aElement0(X5) | ~ aElementOf0(X6, xP)), inference(resolution, [], [f153, f232])).
fof(f232, plain, ! [X0] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xP)), inference(cnf_transformation, [], [f135])).
fof(f153, plain, ! [X2, X0, X1] : (~ sdtlseqdt0(X1, X2) | sdtlseqdt0(X0, X2) | ~ sdtlseqdt0(X0, X1) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f47])).
fof(f47, plain, ! [X0, X1, X2] : (sdtlseqdt0(X0, X2) | ~ sdtlseqdt0(X1, X2) | ~ sdtlseqdt0(X0, X1) | ~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0)), inference(flattening, [], [f46])).
fof(f46, plain, ! [X0, X1, X2] : ((sdtlseqdt0(X0, X2) | (~ sdtlseqdt0(X1, X2) | ~ sdtlseqdt0(X0, X1))) | (~ aElement0(X2) | ~ aElement0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1, X2] : ((aElement0(X2) & aElement0(X1) & aElement0(X0)) => ((sdtlseqdt0(X1, X2) & sdtlseqdt0(X0, X1)) => sdtlseqdt0(X0, X2))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT386+4.p', mTrans)).
fof(f276, plain, (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), sK21) | spl23_3), inference(avatar_component_clause, [], [f274])).
fof(f274, plain, (spl23_3 <=> sdtlseqdt0(sdtlpdtrp0(xf, xp), sK21)), introduced(avatar_definition, [new_symbols(naming, [spl23_3])])).
fof(f26879, plain, (spl23_414 | ~ spl23_4), inference(avatar_split_clause, [], [f26878, f279, f11990])).
fof(f26878, plain, (aElement0(sK21) | ~ spl23_4), inference(subsumption_resolution, [], [f26836, f230])).
fof(f230, plain, aSet0(xP), inference(cnf_transformation, [], [f135])).
fof(f26836, plain, (aElement0(sK21) | ~ aSet0(xP) | ~ spl23_4), inference(resolution, [], [f281, f144])).
fof(f26850, plain, (spl23_412 | ~ spl23_4), inference(avatar_split_clause, [], [f26830, f279, f11956])).
fof(f26830, plain, (aElementOf0(sK21, xU) | ~ spl23_4), inference(resolution, [], [f281, f231])).
fof(f11848, plain, (spl23_410 | ~ spl23_40), inference(avatar_split_clause, [], [f1607, f917, f11843])).
fof(f917, plain, (spl23_40 <=> aElementOf0(sdtlpdtrp0(xf, xp), xU)), introduced(avatar_definition, [new_symbols(naming, [spl23_40])])).
fof(f1607, plain, (aElement0(sdtlpdtrp0(xf, xp)) | ~ spl23_40), inference(subsumption_resolution, [], [f1597, f209])).
fof(f1597, plain, (aElement0(sdtlpdtrp0(xf, xp)) | ~ aSet0(xU) | ~ spl23_40), inference(resolution, [], [f918, f144])).
fof(f918, plain, (aElementOf0(sdtlpdtrp0(xf, xp), xU) | ~ spl23_40), inference(avatar_component_clause, [], [f917])).
fof(f4800, plain, spl23_39, inference(avatar_split_clause, [], [f4799, f913])).
fof(f4799, plain, aSubsetOf0(xT, xU), inference(subsumption_resolution, [], [f4798, f227])).
fof(f227, plain, aSet0(xT), inference(cnf_transformation, [], [f68])).
fof(f4798, plain, (aSubsetOf0(xT, xU) | ~ aSet0(xT)), inference(subsumption_resolution, [], [f4793, f209])).
fof(f4793, plain, (aSubsetOf0(xT, xU) | ~ aSet0(xU) | ~ aSet0(xT)), inference(duplicate_literal_removal, [], [f4789])).
fof(f4789, plain, (aSubsetOf0(xT, xU) | ~ aSet0(xU) | aSubsetOf0(xT, xU) | ~ aSet0(xT) | ~ aSet0(xU)), inference(resolution, [], [f1551, f150])).
fof(f150, plain, ! [X0, X1] : (~ aElementOf0(sK5(X0, X1), X0) | aSubsetOf0(X1, X0) | ~ aSet0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f87])).
fof(f1551, plain, ! [X1] : (aElementOf0(sK5(X1, xT), xU) | aSubsetOf0(xT, X1) | ~ aSet0(X1)), inference(resolution, [], [f524, f263])).
fof(f263, plain, ! [X0] : (~ aElementOf0(X0, xS) | aElementOf0(X0, xU)), inference(forward_demodulation, [], [f221, f259])).
fof(f221, plain, ! [X0] : (aElementOf0(X0, szDzozmdt0(xf)) | ~ aElementOf0(X0, xS)), inference(cnf_transformation, [], [f67])).
fof(f524, plain, ! [X9] : (aElementOf0(sK5(X9, xT), xS) | ~ aSet0(X9) | aSubsetOf0(xT, X9)), inference(subsumption_resolution, [], [f515, f227])).
fof(f515, plain, ! [X9] : (aSubsetOf0(xT, X9) | ~ aSet0(xT) | ~ aSet0(X9) | aElementOf0(sK5(X9, xT), xS)), inference(resolution, [], [f149, f228])).
fof(f149, plain, ! [X0, X1] : (aElementOf0(sK5(X0, X1), X1) | aSubsetOf0(X1, X0) | ~ aSet0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f87])).
fof(f4448, plain, (spl23_80 | ~ spl23_7), inference(avatar_split_clause, [], [f4447, f294, f1437])).
fof(f4447, plain, (sdtlseqdt0(sK22, xp) | ~ spl23_7), inference(subsumption_resolution, [], [f4437, f2227])).
fof(f2227, plain, (aElementOf0(sK22, xU) | ~ spl23_7), inference(resolution, [], [f2218, f263])).
fof(f2218, plain, (aElementOf0(sK22, xS) | ~ spl23_7), inference(resolution, [], [f296, f228])).
fof(f4437, plain, (sdtlseqdt0(sK22, xp) | ~ aElementOf0(sK22, xU) | ~ spl23_7), inference(duplicate_literal_removal, [], [f4433])).
fof(f4433, plain, (sdtlseqdt0(sK22, xp) | ~ aElementOf0(sK22, xU) | sdtlseqdt0(sK22, xp) | ~ aElementOf0(sK22, xU) | ~ spl23_7), inference(resolution, [], [f2243, f244])).
fof(f244, plain, ! [X0] : (~ sdtlseqdt0(X0, sK20(X0)) | sdtlseqdt0(X0, xp) | ~ aElementOf0(X0, xU)), inference(cnf_transformation, [], [f137])).
fof(f2243, plain, (! [X10] : (sdtlseqdt0(sK22, sK20(X10)) | sdtlseqdt0(X10, xp) | ~ aElementOf0(X10, xU)) | ~ spl23_7), inference(resolution, [], [f2217, f243])).
fof(f243, plain, ! [X0] : (aElementOf0(sK20(X0), xP) | sdtlseqdt0(X0, xp) | ~ aElementOf0(X0, xU)), inference(cnf_transformation, [], [f137])).
fof(f2217, plain, (! [X0] : (~ aElementOf0(X0, xP) | sdtlseqdt0(sK22, X0)) | ~ spl23_7), inference(resolution, [], [f296, f233])).
fof(f233, plain, ! [X2, X0] : (~ aElementOf0(X2, xT) | sdtlseqdt0(X2, X0) | ~ aElementOf0(X0, xP)), inference(cnf_transformation, [], [f135])).
fof(f1589, plain, spl23_40, inference(avatar_contradiction_clause, [], [f1588])).
fof(f1588, plain, ($false | spl23_40), inference(subsumption_resolution, [], [f1587, f240])).
fof(f1587, plain, (~ aElementOf0(xp, xU) | spl23_40), inference(resolution, [], [f919, f573])).
fof(f919, plain, (~ aElementOf0(sdtlpdtrp0(xf, xp), xU) | spl23_40), inference(avatar_component_clause, [], [f917])).
fof(f297, plain, (spl23_1 | spl23_7), inference(avatar_split_clause, [], [f250, f294, f265])).
fof(f265, plain, (spl23_1 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl23_1])])).
fof(f250, plain, (aElementOf0(sK22, xT) | sP3), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ((~ aUpperBoundOfIn0(sdtlpdtrp0(xf, xp), xT, xU) & (~ sdtlseqdt0(sK22, sdtlpdtrp0(xf, xp)) & aElementOf0(sK22, xT))) | sP3), inference(skolemisation, [status(esa), new_symbols(skolem, [sK22])], [f78, f142])).
fof(f142, plain, (? [X0] : (~ sdtlseqdt0(X0, sdtlpdtrp0(xf, xp)) & aElementOf0(X0, xT)) => (~ sdtlseqdt0(sK22, sdtlpdtrp0(xf, xp)) & aElementOf0(sK22, xT))), introduced(choice_axiom, [])).
fof(f78, plain, ((~ aUpperBoundOfIn0(sdtlpdtrp0(xf, xp), xT, xU) & ? [X0] : (~ sdtlseqdt0(X0, sdtlpdtrp0(xf, xp)) & aElementOf0(X0, xT))) | sP3), inference(definition_folding, [], [f72, e77])).
fof(f77, plain, ((~ aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) & ? [X1] : (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), X1) & aElementOf0(X1, xP))) | ~ sP3), inference(usedef, [], [e77])).
fof(e77, plain, (sP3 <=> (~ aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) & ? [X1] : (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), X1) & aElementOf0(X1, xP)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f72, plain, ((~ aUpperBoundOfIn0(sdtlpdtrp0(xf, xp), xT, xU) & ? [X0] : (~ sdtlseqdt0(X0, sdtlpdtrp0(xf, xp)) & aElementOf0(X0, xT))) | (~ aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) & ? [X1] : (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), X1) & aElementOf0(X1, xP)))), inference(ennf_transformation, [], [f38])).
fof(f38, plain, ~ ((aUpperBoundOfIn0(sdtlpdtrp0(xf, xp), xT, xU) | ! [X0] : (aElementOf0(X0, xT) => sdtlseqdt0(X0, sdtlpdtrp0(xf, xp)))) & (aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) | ! [X1] : (aElementOf0(X1, xP) => sdtlseqdt0(sdtlpdtrp0(xf, xp), X1)))), inference(rectify, [], [f30])).
fof(f30, plain, ~ ((aUpperBoundOfIn0(sdtlpdtrp0(xf, xp), xT, xU) | ! [X0] : (aElementOf0(X0, xT) => sdtlseqdt0(X0, sdtlpdtrp0(xf, xp)))) & (aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) | ! [X0] : (aElementOf0(X0, xP) => sdtlseqdt0(sdtlpdtrp0(xf, xp), X0)))), inference(negated_conjecture, [], [f29])).
fof(f29, plain, ~ ((aUpperBoundOfIn0(sdtlpdtrp0(xf, xp), xT, xU) | ! [X0] : (aElementOf0(X0, xT) => sdtlseqdt0(X0, sdtlpdtrp0(xf, xp)))) & (aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) | ! [X0] : (aElementOf0(X0, xP) => sdtlseqdt0(sdtlpdtrp0(xf, xp), X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT386+4.p', m__)).
fof(f292, plain, (spl23_1 | ~ spl23_6), inference(avatar_split_clause, [], [f251, f289, f265])).
fof(f251, plain, (~ sdtlseqdt0(sK22, sdtlpdtrp0(xf, xp)) | sP3), inference(cnf_transformation, [], [f143])).
fof(f282, plain, (~ spl23_1 | spl23_4), inference(avatar_split_clause, [], [f247, f279, f265])).
fof(f247, plain, (aElementOf0(sK21, xP) | ~ sP3), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ((~ aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) & (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), sK21) & aElementOf0(sK21, xP))) | ~ sP3), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21])], [f139, f140])).
fof(f140, plain, (? [X0] : (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), X0) & aElementOf0(X0, xP)) => (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), sK21) & aElementOf0(sK21, xP))), introduced(choice_axiom, [])).
fof(f139, plain, ((~ aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) & ? [X0] : (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), X0) & aElementOf0(X0, xP))) | ~ sP3), inference(rectify, [], [f138])).
fof(f138, plain, ((~ aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) & ? [X1] : (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), X1) & aElementOf0(X1, xP))) | ~ sP3), inference(nnf_transformation, [], [f77])).
fof(f277, plain, (~ spl23_1 | ~ spl23_3), inference(avatar_split_clause, [], [f248, f274, f265])).
fof(f248, plain, (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), sK21) | ~ sP3), inference(cnf_transformation, [], [f141])).