fof(f428001, plain, $false, inference(avatar_sat_refutation, [], [f473, f675, f1291, f1326, f1338, f1365, f1391, f1416, f1421, f1430, f1438, f1656, f1659, f2111, f2137, f11121, f11150, f11816, f12151, f12166, f12169, f13010, f13029, f162084, f169377, f221273, f221406, f221505, f221803, f221936, f222035, f427987, f428000])).
fof(f428000, plain, (~ spl28_1 | spl28_2 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183), inference(avatar_contradiction_clause, [], [f427999])).
fof(f427999, plain, ($false | (~ spl28_1 | spl28_2 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(subsumption_resolution, [], [f427998, f427995])).
fof(f427995, plain, (~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), k19_filter_2(sK1, k2_xboole_0(sK2, sK3))) | (spl28_2 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154)), inference(forward_demodulation, [], [f427994, f2173])).
fof(f2173, plain, (k4_subset_1(u1_struct_0(sK1), sK2, sK3) = k2_xboole_0(sK2, sK3)), inference(resolution, [], [f836, f299])).
fof(f299, plain, m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))), inference(cnf_transformation, [], [f242])).
fof(f242, plain, ((((~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, k19_filter_2(sK1, sK3)))) | ~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3)))) & m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) & ~ v1_xboole_0(sK3)) & m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) & ~ v1_xboole_0(sK2)) & l3_lattices(sK1) & v10_lattices(sK1) & ~ v3_struct_0(sK1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2, sK3])], [f118, f241, f240, f239])).
fof(f239, plain, (? [X0] : (? [X1] : (? [X2] : ((~ r1_filter_2(u1_struct_0(X0), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, X2)), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, k19_filter_2(X0, X2)))) | ~ r1_filter_2(u1_struct_0(X0), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, X2)), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), k19_filter_2(X0, X1), X2)))) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X2)) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (? [X1] : (? [X2] : ((~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X1, X2)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X1, k19_filter_2(sK1, X2)))) | ~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X1, X2)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, X1), X2)))) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK1))) & ~ v1_xboole_0(X2)) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(sK1))) & ~ v1_xboole_0(X1)) & l3_lattices(sK1) & v10_lattices(sK1) & ~ v3_struct_0(sK1))), introduced(choice_axiom, [])).
fof(f240, plain, (? [X1] : (? [X2] : ((~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X1, X2)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X1, k19_filter_2(sK1, X2)))) | ~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X1, X2)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, X1), X2)))) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK1))) & ~ v1_xboole_0(X2)) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(sK1))) & ~ v1_xboole_0(X1)) => (? [X2] : ((~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, X2)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, k19_filter_2(sK1, X2)))) | ~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, X2)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), X2)))) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK1))) & ~ v1_xboole_0(X2)) & m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) & ~ v1_xboole_0(sK2))), introduced(choice_axiom, [])).
fof(f241, plain, (? [X2] : ((~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, X2)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, k19_filter_2(sK1, X2)))) | ~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, X2)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), X2)))) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK1))) & ~ v1_xboole_0(X2)) => ((~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, k19_filter_2(sK1, sK3)))) | ~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3)))) & m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) & ~ v1_xboole_0(sK3))), introduced(choice_axiom, [])).
fof(f118, plain, ? [X0] : (? [X1] : (? [X2] : ((~ r1_filter_2(u1_struct_0(X0), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, X2)), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, k19_filter_2(X0, X2)))) | ~ r1_filter_2(u1_struct_0(X0), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, X2)), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), k19_filter_2(X0, X1), X2)))) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X2)) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f117])).
fof(f117, plain, ? [X0] : (? [X1] : (? [X2] : ((~ r1_filter_2(u1_struct_0(X0), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, X2)), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, k19_filter_2(X0, X2)))) | ~ r1_filter_2(u1_struct_0(X0), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, X2)), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), k19_filter_2(X0, X1), X2)))) & (m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X2))) & (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1))) & (l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) => ! [X2] : ((m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X2)) => (r1_filter_2(u1_struct_0(X0), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, X2)), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, k19_filter_2(X0, X2)))) & r1_filter_2(u1_struct_0(X0), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, X2)), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), k19_filter_2(X0, X1), X2))))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) => ! [X2] : ((m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X2)) => (r1_filter_2(u1_struct_0(X0), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, X2)), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, k19_filter_2(X0, X2)))) & r1_filter_2(u1_struct_0(X0), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), X1, X2)), k19_filter_2(X0, k4_subset_1(u1_struct_0(X0), k19_filter_2(X0, X1), X2))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', t49_filter_2)).
fof(f836, plain, ! [X20] : (~ m1_subset_1(X20, k1_zfmisc_1(u1_struct_0(sK1))) | (k4_subset_1(u1_struct_0(sK1), X20, sK3) = k2_xboole_0(X20, sK3))), inference(resolution, [], [f438, f301])).
fof(f301, plain, m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))), inference(cnf_transformation, [], [f242])).
fof(f438, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, k1_zfmisc_1(X0)) | (k4_subset_1(X0, X1, X2) = k2_xboole_0(X1, X2)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0))), inference(cnf_transformation, [], [f214])).
fof(f214, plain, ! [X0, X1, X2] : ((k4_subset_1(X0, X1, X2) = k2_xboole_0(X1, X2)) | ~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0))), inference(flattening, [], [f213])).
fof(f213, plain, ! [X0, X1, X2] : ((k4_subset_1(X0, X1, X2) = k2_xboole_0(X1, X2)) | (~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0)))), inference(ennf_transformation, [], [f86])).
fof(f86, plain, ! [X0, X1, X2] : ((m1_subset_1(X2, k1_zfmisc_1(X0)) & m1_subset_1(X1, k1_zfmisc_1(X0))) => (k4_subset_1(X0, X1, X2) = k2_xboole_0(X1, X2))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', redefinition_k4_subset_1)).
fof(f427994, plain, (~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k2_xboole_0(sK2, sK3))) | (spl28_2 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154)), inference(forward_demodulation, [], [f427993, f427898])).
fof(f427898, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154)), inference(backward_demodulation, [], [f365861, f427897])).
fof(f427897, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (~ spl28_19 | ~ spl28_24 | ~ spl28_32 | ~ spl28_69 | ~ spl28_285)), inference(forward_demodulation, [], [f427896, f11796])).
fof(f11796, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, sK3))) | ~ spl28_285), inference(avatar_component_clause, [], [f11794])).
fof(f11794, plain, (spl28_285 <=> (k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, sK3)))), introduced(avatar_definition, [new_symbols(naming, [spl28_285])])).
fof(f427896, plain, ((k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, sK3)) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (~ spl28_19 | ~ spl28_24 | ~ spl28_32 | ~ spl28_69)), inference(forward_demodulation, [], [f427895, f15480])).
fof(f15480, plain, (k2_xboole_0(sK3, k19_filter_2(sK1, sK2)) = k4_subset_1(u1_struct_0(sK1), sK3, k19_filter_2(sK1, sK2))), inference(resolution, [], [f5148, f954])).
fof(f954, plain, m2_lattice4(k19_filter_2(sK1, sK2), sK1), inference(subsumption_resolution, [], [f953, f295])).
fof(f295, plain, ~ v3_struct_0(sK1), inference(cnf_transformation, [], [f242])).
fof(f953, plain, (m2_lattice4(k19_filter_2(sK1, sK2), sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f952, f296])).
fof(f296, plain, v10_lattices(sK1), inference(cnf_transformation, [], [f242])).
fof(f952, plain, (m2_lattice4(k19_filter_2(sK1, sK2), sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f947, f297])).
fof(f297, plain, l3_lattices(sK1), inference(cnf_transformation, [], [f242])).
fof(f947, plain, (m2_lattice4(k19_filter_2(sK1, sK2), sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f939, f343])).
fof(f343, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f161])).
fof(f161, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_m2_filter_2)).
fof(f939, plain, m2_filter_2(k19_filter_2(sK1, sK2), sK1), inference(subsumption_resolution, [], [f938, f295])).
fof(f938, plain, (m2_filter_2(k19_filter_2(sK1, sK2), sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f937, f296])).
fof(f937, plain, (m2_filter_2(k19_filter_2(sK1, sK2), sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f936, f297])).
fof(f936, plain, (m2_filter_2(k19_filter_2(sK1, sK2), sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f927, f298])).
fof(f298, plain, ~ v1_xboole_0(sK2), inference(cnf_transformation, [], [f242])).
fof(f927, plain, (m2_filter_2(k19_filter_2(sK1, sK2), sK1) | v1_xboole_0(sK2) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f326, f299])).
fof(f326, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | m2_filter_2(k19_filter_2(X0, X1), X0) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ! [X0, X1] : (m2_filter_2(k19_filter_2(X0, X1), X0) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f141])).
fof(f141, plain, ! [X0, X1] : (m2_filter_2(k19_filter_2(X0, X1), X0) | (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0, X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m2_filter_2(k19_filter_2(X0, X1), X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_k19_filter_2)).
fof(f5148, plain, ! [X24] : (~ m2_lattice4(X24, sK1) | (k4_subset_1(u1_struct_0(sK1), sK3, X24) = k2_xboole_0(sK3, X24))), inference(subsumption_resolution, [], [f5147, f295])).
fof(f5147, plain, ! [X24] : ((k4_subset_1(u1_struct_0(sK1), sK3, X24) = k2_xboole_0(sK3, X24)) | ~ m2_lattice4(X24, sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f5146, f296])).
fof(f5146, plain, ! [X24] : ((k4_subset_1(u1_struct_0(sK1), sK3, X24) = k2_xboole_0(sK3, X24)) | ~ m2_lattice4(X24, sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f5119, f297])).
fof(f5119, plain, ! [X24] : ((k4_subset_1(u1_struct_0(sK1), sK3, X24) = k2_xboole_0(sK3, X24)) | ~ m2_lattice4(X24, sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f831, f301])).
fof(f831, plain, ! [X8, X7, X9] : (~ m1_subset_1(X8, k1_zfmisc_1(u1_struct_0(X7))) | (k4_subset_1(u1_struct_0(X7), X8, X9) = k2_xboole_0(X8, X9)) | ~ m2_lattice4(X9, X7) | ~ l3_lattices(X7) | ~ v10_lattices(X7) | v3_struct_0(X7)), inference(resolution, [], [f438, f344])).
fof(f344, plain, ! [X0, X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f163])).
fof(f163, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_lattice4(X1, X0) => m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_m2_lattice4)).
fof(f427895, plain, ((k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, sK3)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK3, k19_filter_2(sK1, sK2)))) | (~ spl28_19 | ~ spl28_24 | ~ spl28_32 | ~ spl28_69)), inference(forward_demodulation, [], [f427894, f2078])).
fof(f2078, plain, (k4_subset_1(u1_struct_0(sK1), sK3, sK2) = k2_xboole_0(sK2, sK3)), inference(forward_demodulation, [], [f2060, f318])).
fof(f318, plain, ! [X0, X1] : (k2_xboole_0(X0, X1) = k2_xboole_0(X1, X0)), inference(cnf_transformation, [], [f10])).
fof(f10, plain, ! [X0, X1] : (k2_xboole_0(X0, X1) = k2_xboole_0(X1, X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', commutativity_k2_xboole_0)).
fof(f2060, plain, (k4_subset_1(u1_struct_0(sK1), sK3, sK2) = k2_xboole_0(sK3, sK2)), inference(resolution, [], [f835, f301])).
fof(f835, plain, ! [X19] : (~ m1_subset_1(X19, k1_zfmisc_1(u1_struct_0(sK1))) | (k4_subset_1(u1_struct_0(sK1), X19, sK2) = k2_xboole_0(X19, sK2))), inference(resolution, [], [f438, f299])).
fof(f427894, plain, ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK3, k19_filter_2(sK1, sK2))) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK3, sK2))) | (~ spl28_19 | ~ spl28_24 | ~ spl28_32 | ~ spl28_69)), inference(subsumption_resolution, [], [f427727, f300])).
fof(f300, plain, ~ v1_xboole_0(sK3), inference(cnf_transformation, [], [f242])).
fof(f427727, plain, ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK3, k19_filter_2(sK1, sK2))) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK3, sK2))) | v1_xboole_0(sK3) | (~ spl28_19 | ~ spl28_24 | ~ spl28_32 | ~ spl28_69)), inference(resolution, [], [f16224, f301])).
fof(f16224, plain, (! [X22] : (~ m1_subset_1(X22, k1_zfmisc_1(u1_struct_0(sK1))) | (k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), X22, sK2)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), X22, k19_filter_2(sK1, sK2)))) | v1_xboole_0(X22)) | (~ spl28_19 | ~ spl28_24 | ~ spl28_32 | ~ spl28_69)), inference(forward_demodulation, [], [f16064, f15995])).
fof(f15995, plain, ((u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | ~ spl28_69), inference(trivial_inequality_removal, [], [f15981])).
fof(f15981, plain, (~ (k1_lattice2(sK1) = k1_lattice2(sK1)) | (u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | ~ spl28_69), inference(superposition, [], [f1655, f2485])).
fof(f2485, plain, (k1_lattice2(sK1) = g3_lattices(u1_struct_0(k1_lattice2(sK1)), u2_lattices(k1_lattice2(sK1)), u1_lattices(k1_lattice2(sK1)))), inference(resolution, [], [f755, f297])).
fof(f755, plain, ! [X1] : (~ l3_lattices(X1) | (k1_lattice2(X1) = g3_lattices(u1_struct_0(k1_lattice2(X1)), u2_lattices(k1_lattice2(X1)), u1_lattices(k1_lattice2(X1))))), inference(subsumption_resolution, [], [f751, f328])).
fof(f328, plain, ! [X0] : (l3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ! [X0] : ((l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0] : (l3_lattices(X0) => (l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_k1_lattice2)).
fof(f751, plain, ! [X1] : ((k1_lattice2(X1) = g3_lattices(u1_struct_0(k1_lattice2(X1)), u2_lattices(k1_lattice2(X1)), u1_lattices(k1_lattice2(X1)))) | ~ l3_lattices(k1_lattice2(X1)) | ~ l3_lattices(X1)), inference(resolution, [], [f303, f327])).
fof(f327, plain, ! [X0] : (v3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f143])).
fof(f303, plain, ! [X0] : (~ v3_lattices(X0) | (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ! [X0] : ((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f119])).
fof(f119, plain, ! [X0] : (((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (l3_lattices(X0) => (v3_lattices(X0) => (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', abstractness_v3_lattices)).
fof(f1655, plain, (! [X17, X15, X16] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15)) | ~ spl28_69), inference(avatar_component_clause, [], [f1654])).
fof(f1654, plain, (spl28_69 <=> ! [X16, X15, X17] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15))), introduced(avatar_definition, [new_symbols(naming, [spl28_69])])).
fof(f16064, plain, (! [X22] : (~ m1_subset_1(X22, k1_zfmisc_1(u1_struct_0(sK1))) | (k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, sK2)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, k19_filter_2(sK1, sK2)))) | v1_xboole_0(X22)) | (~ spl28_19 | ~ spl28_24 | ~ spl28_32 | ~ spl28_69)), inference(backward_demodulation, [], [f2161, f15995])).
fof(f2161, plain, (! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, sK2)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, k19_filter_2(sK1, sK2)))) | v1_xboole_0(X22) | ~ m1_subset_1(X22, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1))))) | (~ spl28_19 | ~ spl28_24 | ~ spl28_32)), inference(backward_demodulation, [], [f1420, f2158])).
fof(f2158, plain, ((k19_filter_2(sK1, sK2) = k3_filter_0(k1_lattice2(sK1), sK2)) | (~ spl28_19 | ~ spl28_24)), inference(backward_demodulation, [], [f1325, f2157])).
fof(f2157, plain, ((k19_filter_2(sK1, sK2) = k3_filter_2(k1_lattice2(sK1), sK2)) | ~ spl28_19), inference(forward_demodulation, [], [f1290, f1307])).
fof(f1307, plain, (sK2 = k8_filter_2(sK1, sK2)), inference(subsumption_resolution, [], [f1306, f295])).
fof(f1306, plain, ((sK2 = k8_filter_2(sK1, sK2)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1305, f296])).
fof(f1305, plain, ((sK2 = k8_filter_2(sK1, sK2)) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1266, f297])).
fof(f1266, plain, ((sK2 = k8_filter_2(sK1, sK2)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f1081, f322])).
fof(f322, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) | (k8_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ! [X0] : (! [X1] : ((k8_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0))))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f135])).
fof(f135, plain, ! [X0] : (! [X1] : ((k8_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0))))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) => (k8_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', d7_filter_2)).
fof(f1081, plain, m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))), inference(subsumption_resolution, [], [f1080, f295])).
fof(f1080, plain, (m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1079, f296])).
fof(f1079, plain, (m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1078, f297])).
fof(f1078, plain, (m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1063, f299])).
fof(f1063, plain, (m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(superposition, [], [f332, f823])).
fof(f823, plain, (sK2 = k7_filter_2(sK1, sK2)), inference(subsumption_resolution, [], [f822, f295])).
fof(f822, plain, ((sK2 = k7_filter_2(sK1, sK2)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f821, f296])).
fof(f821, plain, ((sK2 = k7_filter_2(sK1, sK2)) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f813, f297])).
fof(f813, plain, ((sK2 = k7_filter_2(sK1, sK2)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f321, f299])).
fof(f321, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | (k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f133])).
fof(f133, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) => (k7_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', d6_filter_2)).
fof(f332, plain, ! [X0, X1] : (m1_subset_1(k7_filter_2(X0, X1), k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ! [X0, X1] : (m1_subset_1(k7_filter_2(X0, X1), k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f150])).
fof(f150, plain, ! [X0, X1] : (m1_subset_1(k7_filter_2(X0, X1), k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) | (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ! [X0, X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_subset_1(k7_filter_2(X0, X1), k1_zfmisc_1(u1_struct_0(k1_lattice2(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_k7_filter_2)).
fof(f1290, plain, ((k3_filter_2(k1_lattice2(sK1), sK2) = k19_filter_2(sK1, k8_filter_2(sK1, sK2))) | ~ spl28_19), inference(avatar_component_clause, [], [f1288])).
fof(f1288, plain, (spl28_19 <=> (k3_filter_2(k1_lattice2(sK1), sK2) = k19_filter_2(sK1, k8_filter_2(sK1, sK2)))), introduced(avatar_definition, [new_symbols(naming, [spl28_19])])).
fof(f1325, plain, ((k3_filter_2(k1_lattice2(sK1), sK2) = k3_filter_0(k1_lattice2(sK1), sK2)) | ~ spl28_24), inference(avatar_component_clause, [], [f1323])).
fof(f1323, plain, (spl28_24 <=> (k3_filter_2(k1_lattice2(sK1), sK2) = k3_filter_0(k1_lattice2(sK1), sK2))), introduced(avatar_definition, [new_symbols(naming, [spl28_24])])).
fof(f1420, plain, (! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, sK2)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, k3_filter_0(k1_lattice2(sK1), sK2)))) | v1_xboole_0(X22) | ~ m1_subset_1(X22, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1))))) | ~ spl28_32), inference(avatar_component_clause, [], [f1419])).
fof(f1419, plain, (spl28_32 <=> ! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, sK2)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, k3_filter_0(k1_lattice2(sK1), sK2)))) | v1_xboole_0(X22) | ~ m1_subset_1(X22, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))))), introduced(avatar_definition, [new_symbols(naming, [spl28_32])])).
fof(f365861, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_69 | ~ spl28_80 | ~ spl28_2141 | ~ spl28_2154)), inference(forward_demodulation, [], [f365860, f365671])).
fof(f365671, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(k19_filter_2(sK1, sK2), k19_filter_2(sK1, sK3)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_69 | ~ spl28_80 | ~ spl28_2141 | ~ spl28_2154)), inference(forward_demodulation, [], [f365670, f221538])).
fof(f221538, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | (~ spl28_2141 | ~ spl28_2154)), inference(backward_demodulation, [], [f221405, f221272])).
fof(f221272, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | ~ spl28_2141), inference(avatar_component_clause, [], [f221270])).
fof(f221270, plain, (spl28_2141 <=> (k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))), introduced(avatar_definition, [new_symbols(naming, [spl28_2141])])).
fof(f221405, plain, ((k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | ~ spl28_2154), inference(avatar_component_clause, [], [f221403])).
fof(f221403, plain, (spl28_2154 <=> (k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))), introduced(avatar_definition, [new_symbols(naming, [spl28_2154])])).
fof(f365670, plain, ((k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(k19_filter_2(sK1, sK2), k19_filter_2(sK1, sK3)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_69 | ~ spl28_80)), inference(forward_demodulation, [], [f365669, f15451])).
fof(f15451, plain, (k4_subset_1(u1_struct_0(sK1), sK2, k19_filter_2(sK1, sK3)) = k2_xboole_0(sK2, k19_filter_2(sK1, sK3))), inference(resolution, [], [f5139, f1014])).
fof(f1014, plain, m2_lattice4(k19_filter_2(sK1, sK3), sK1), inference(subsumption_resolution, [], [f1013, f295])).
fof(f1013, plain, (m2_lattice4(k19_filter_2(sK1, sK3), sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1012, f296])).
fof(f1012, plain, (m2_lattice4(k19_filter_2(sK1, sK3), sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1007, f297])).
fof(f1007, plain, (m2_lattice4(k19_filter_2(sK1, sK3), sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f943, f343])).
fof(f943, plain, m2_filter_2(k19_filter_2(sK1, sK3), sK1), inference(subsumption_resolution, [], [f942, f295])).
fof(f942, plain, (m2_filter_2(k19_filter_2(sK1, sK3), sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f941, f296])).
fof(f941, plain, (m2_filter_2(k19_filter_2(sK1, sK3), sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f940, f297])).
fof(f940, plain, (m2_filter_2(k19_filter_2(sK1, sK3), sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f928, f300])).
fof(f928, plain, (m2_filter_2(k19_filter_2(sK1, sK3), sK1) | v1_xboole_0(sK3) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f326, f301])).
fof(f5139, plain, ! [X21] : (~ m2_lattice4(X21, sK1) | (k4_subset_1(u1_struct_0(sK1), sK2, X21) = k2_xboole_0(sK2, X21))), inference(subsumption_resolution, [], [f5138, f295])).
fof(f5138, plain, ! [X21] : ((k4_subset_1(u1_struct_0(sK1), sK2, X21) = k2_xboole_0(sK2, X21)) | ~ m2_lattice4(X21, sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f5137, f296])).
fof(f5137, plain, ! [X21] : ((k4_subset_1(u1_struct_0(sK1), sK2, X21) = k2_xboole_0(sK2, X21)) | ~ m2_lattice4(X21, sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f5116, f297])).
fof(f5116, plain, ! [X21] : ((k4_subset_1(u1_struct_0(sK1), sK2, X21) = k2_xboole_0(sK2, X21)) | ~ m2_lattice4(X21, sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f831, f299])).
fof(f365669, plain, ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK2, k19_filter_2(sK1, sK3))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(k19_filter_2(sK1, sK2), k19_filter_2(sK1, sK3)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_69 | ~ spl28_80)), inference(forward_demodulation, [], [f365605, f265844])).
fof(f265844, plain, (k2_xboole_0(k19_filter_2(sK1, sK2), k19_filter_2(sK1, sK3)) = k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), k19_filter_2(sK1, sK3))), inference(resolution, [], [f38005, f954])).
fof(f38005, plain, ! [X59] : (~ m2_lattice4(X59, sK1) | (k2_xboole_0(X59, k19_filter_2(sK1, sK3)) = k4_subset_1(u1_struct_0(sK1), X59, k19_filter_2(sK1, sK3)))), inference(subsumption_resolution, [], [f38004, f295])).
fof(f38004, plain, ! [X59] : ((k2_xboole_0(X59, k19_filter_2(sK1, sK3)) = k4_subset_1(u1_struct_0(sK1), X59, k19_filter_2(sK1, sK3))) | v3_struct_0(sK1) | ~ m2_lattice4(X59, sK1)), inference(subsumption_resolution, [], [f38003, f296])).
fof(f38003, plain, ! [X59] : ((k2_xboole_0(X59, k19_filter_2(sK1, sK3)) = k4_subset_1(u1_struct_0(sK1), X59, k19_filter_2(sK1, sK3))) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m2_lattice4(X59, sK1)), inference(subsumption_resolution, [], [f37896, f297])).
fof(f37896, plain, ! [X59] : ((k2_xboole_0(X59, k19_filter_2(sK1, sK3)) = k4_subset_1(u1_struct_0(sK1), X59, k19_filter_2(sK1, sK3))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m2_lattice4(X59, sK1)), inference(resolution, [], [f5133, f1014])).
fof(f5133, plain, ! [X2, X0, X1] : (~ m2_lattice4(X2, X0) | (k4_subset_1(u1_struct_0(X0), X1, X2) = k2_xboole_0(X1, X2)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_lattice4(X1, X0)), inference(duplicate_literal_removal, [], [f5109])).
fof(f5109, plain, ! [X2, X0, X1] : ((k4_subset_1(u1_struct_0(X0), X1, X2) = k2_xboole_0(X1, X2)) | ~ m2_lattice4(X2, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f831, f344])).
fof(f365605, plain, ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK2, k19_filter_2(sK1, sK3))) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), k19_filter_2(sK1, sK3)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_69 | ~ spl28_80)), inference(resolution, [], [f16141, f2165])).
fof(f2165, plain, (m1_filter_0(k19_filter_2(sK1, sK3), k1_lattice2(sK1)) | (spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_80)), inference(backward_demodulation, [], [f1780, f2136])).
fof(f2136, plain, ((k19_filter_2(sK1, sK3) = k3_filter_0(k1_lattice2(sK1), sK3)) | ~ spl28_80), inference(avatar_component_clause, [], [f2134])).
fof(f2134, plain, (spl28_80 <=> (k19_filter_2(sK1, sK3) = k3_filter_0(k1_lattice2(sK1), sK3))), introduced(avatar_definition, [new_symbols(naming, [spl28_80])])).
fof(f1780, plain, (m1_filter_0(k3_filter_0(k1_lattice2(sK1), sK3), k1_lattice2(sK1)) | (spl28_21 | ~ spl28_22 | ~ spl28_23)), inference(subsumption_resolution, [], [f1779, f1312])).
fof(f1312, plain, (~ v3_struct_0(k1_lattice2(sK1)) | spl28_21), inference(avatar_component_clause, [], [f1311])).
fof(f1311, plain, (spl28_21 <=> v3_struct_0(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_21])])).
fof(f1779, plain, (m1_filter_0(k3_filter_0(k1_lattice2(sK1), sK3), k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl28_22 | ~ spl28_23)), inference(subsumption_resolution, [], [f1778, f1316])).
fof(f1316, plain, (v10_lattices(k1_lattice2(sK1)) | ~ spl28_22), inference(avatar_component_clause, [], [f1315])).
fof(f1315, plain, (spl28_22 <=> v10_lattices(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_22])])).
fof(f1778, plain, (m1_filter_0(k3_filter_0(k1_lattice2(sK1), sK3), k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ spl28_23), inference(subsumption_resolution, [], [f1777, f1320])).
fof(f1320, plain, (l3_lattices(k1_lattice2(sK1)) | ~ spl28_23), inference(avatar_component_clause, [], [f1319])).
fof(f1319, plain, (spl28_23 <=> l3_lattices(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_23])])).
fof(f1777, plain, (m1_filter_0(k3_filter_0(k1_lattice2(sK1), sK3), k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(subsumption_resolution, [], [f1749, f300])).
fof(f1749, plain, (m1_filter_0(k3_filter_0(k1_lattice2(sK1), sK3), k1_lattice2(sK1)) | v1_xboole_0(sK3) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(resolution, [], [f1085, f329])).
fof(f329, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | m1_filter_0(k3_filter_0(X0, X1), X0) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ! [X0, X1] : (m1_filter_0(k3_filter_0(X0, X1), X0) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f144])).
fof(f144, plain, ! [X0, X1] : (m1_filter_0(k3_filter_0(X0, X1), X0) | (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ! [X0, X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_filter_0(k3_filter_0(X0, X1), X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_k3_filter_0)).
fof(f1085, plain, m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))), inference(subsumption_resolution, [], [f1084, f295])).
fof(f1084, plain, (m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1083, f296])).
fof(f1083, plain, (m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1082, f297])).
fof(f1082, plain, (m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1064, f301])).
fof(f1064, plain, (m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(superposition, [], [f332, f826])).
fof(f826, plain, (sK3 = k7_filter_2(sK1, sK3)), inference(subsumption_resolution, [], [f825, f295])).
fof(f825, plain, ((sK3 = k7_filter_2(sK1, sK3)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f824, f296])).
fof(f824, plain, ((sK3 = k7_filter_2(sK1, sK3)) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f814, f297])).
fof(f814, plain, ((sK3 = k7_filter_2(sK1, sK3)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f321, f301])).
fof(f16141, plain, (! [X22] : (~ m1_filter_0(X22, k1_lattice2(sK1)) | (k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK2, X22)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), X22)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_69)), inference(backward_demodulation, [], [f8690, f15995])).
fof(f8690, plain, (! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK2, X22)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k19_filter_2(sK1, sK2), X22))) | ~ m1_filter_0(X22, k1_lattice2(sK1))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24)), inference(forward_demodulation, [], [f8689, f2158])).
fof(f8689, plain, (! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK2, X22)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k3_filter_0(k1_lattice2(sK1), sK2), X22))) | ~ m1_filter_0(X22, k1_lattice2(sK1))) | (spl28_21 | ~ spl28_22 | ~ spl28_23)), inference(subsumption_resolution, [], [f8688, f1312])).
fof(f8688, plain, (! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK2, X22)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k3_filter_0(k1_lattice2(sK1), sK2), X22))) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_filter_0(X22, k1_lattice2(sK1))) | (~ spl28_22 | ~ spl28_23)), inference(subsumption_resolution, [], [f8687, f1316])).
fof(f8687, plain, (! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK2, X22)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k3_filter_0(k1_lattice2(sK1), sK2), X22))) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_filter_0(X22, k1_lattice2(sK1))) | ~ spl28_23), inference(subsumption_resolution, [], [f8686, f1320])).
fof(f8686, plain, ! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK2, X22)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k3_filter_0(k1_lattice2(sK1), sK2), X22))) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_filter_0(X22, k1_lattice2(sK1))), inference(subsumption_resolution, [], [f8662, f298])).
fof(f8662, plain, ! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK2, X22)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k3_filter_0(k1_lattice2(sK1), sK2), X22))) | v1_xboole_0(sK2) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_filter_0(X22, k1_lattice2(sK1))), inference(resolution, [], [f1384, f1081])).
fof(f1384, plain, ! [X4, X5, X3] : (~ m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X3))) | (k3_filter_0(X3, k4_subset_1(u1_struct_0(X3), X4, X5)) = k3_filter_0(X3, k4_subset_1(u1_struct_0(X3), k3_filter_0(X3, X4), X5))) | v1_xboole_0(X4) | ~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | ~ m1_filter_0(X5, X3)), inference(subsumption_resolution, [], [f1382, f338])).
fof(f338, plain, ! [X0, X1] : (~ m1_filter_0(X1, X0) | ~ v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ! [X0] : (! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f157])).
fof(f157, plain, ! [X0] : (! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) | ~ m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_0(X1, X0) => (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_m1_filter_0)).
fof(f1382, plain, ! [X4, X5, X3] : ((k3_filter_0(X3, k4_subset_1(u1_struct_0(X3), X4, X5)) = k3_filter_0(X3, k4_subset_1(u1_struct_0(X3), k3_filter_0(X3, X4), X5))) | v1_xboole_0(X5) | ~ m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X3))) | v1_xboole_0(X4) | ~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | ~ m1_filter_0(X5, X3)), inference(duplicate_literal_removal, [], [f1368])).
fof(f1368, plain, ! [X4, X5, X3] : ((k3_filter_0(X3, k4_subset_1(u1_struct_0(X3), X4, X5)) = k3_filter_0(X3, k4_subset_1(u1_struct_0(X3), k3_filter_0(X3, X4), X5))) | v1_xboole_0(X5) | ~ m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X3))) | v1_xboole_0(X4) | ~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | ~ m1_filter_0(X5, X3) | ~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3)), inference(resolution, [], [f456, f339])).
fof(f339, plain, ! [X0, X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f158])).
fof(f456, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) | (k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, X2)) = k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), k3_filter_0(X0, X1), X2))) | v1_xboole_0(X2) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f230])).
fof(f230, plain, ! [X0] : (! [X1] : (! [X2] : (((k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, X2)) = k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, k3_filter_0(X0, X2)))) & (k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, X2)) = k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), k3_filter_0(X0, X1), X2)))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X2)) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f229])).
fof(f229, plain, ! [X0] : (! [X1] : (! [X2] : (((k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, X2)) = k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, k3_filter_0(X0, X2)))) & (k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, X2)) = k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), k3_filter_0(X0, X1), X2)))) | (~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X2))) | (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) => ! [X2] : ((m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X2)) => ((k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, X2)) = k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, k3_filter_0(X0, X2)))) & (k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, X2)) = k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), k3_filter_0(X0, X1), X2))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', t47_filter_0)).
fof(f365860, plain, ((k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(k19_filter_2(sK1, sK2), k19_filter_2(sK1, sK3)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_69 | ~ spl28_80)), inference(forward_demodulation, [], [f365859, f15480])).
fof(f365859, plain, ((k3_filter_0(k1_lattice2(sK1), k2_xboole_0(k19_filter_2(sK1, sK2), k19_filter_2(sK1, sK3))) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK3, k19_filter_2(sK1, sK2)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_69 | ~ spl28_80)), inference(forward_demodulation, [], [f365817, f265689])).
fof(f265689, plain, (k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK3), k19_filter_2(sK1, sK2)) = k2_xboole_0(k19_filter_2(sK1, sK2), k19_filter_2(sK1, sK3))), inference(forward_demodulation, [], [f265609, f318])).
fof(f265609, plain, (k2_xboole_0(k19_filter_2(sK1, sK3), k19_filter_2(sK1, sK2)) = k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK3), k19_filter_2(sK1, sK2))), inference(resolution, [], [f38002, f1014])).
fof(f38002, plain, ! [X58] : (~ m2_lattice4(X58, sK1) | (k2_xboole_0(X58, k19_filter_2(sK1, sK2)) = k4_subset_1(u1_struct_0(sK1), X58, k19_filter_2(sK1, sK2)))), inference(subsumption_resolution, [], [f38001, f295])).
fof(f38001, plain, ! [X58] : ((k2_xboole_0(X58, k19_filter_2(sK1, sK2)) = k4_subset_1(u1_struct_0(sK1), X58, k19_filter_2(sK1, sK2))) | v3_struct_0(sK1) | ~ m2_lattice4(X58, sK1)), inference(subsumption_resolution, [], [f38000, f296])).
fof(f38000, plain, ! [X58] : ((k2_xboole_0(X58, k19_filter_2(sK1, sK2)) = k4_subset_1(u1_struct_0(sK1), X58, k19_filter_2(sK1, sK2))) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m2_lattice4(X58, sK1)), inference(subsumption_resolution, [], [f37895, f297])).
fof(f37895, plain, ! [X58] : ((k2_xboole_0(X58, k19_filter_2(sK1, sK2)) = k4_subset_1(u1_struct_0(sK1), X58, k19_filter_2(sK1, sK2))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m2_lattice4(X58, sK1)), inference(resolution, [], [f5133, f954])).
fof(f365817, plain, ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK3, k19_filter_2(sK1, sK2))) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK3), k19_filter_2(sK1, sK2)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_69 | ~ spl28_80)), inference(resolution, [], [f16142, f2160])).
fof(f2160, plain, (m1_filter_0(k19_filter_2(sK1, sK2), k1_lattice2(sK1)) | (~ spl28_19 | ~ spl28_24 | ~ spl28_26)), inference(backward_demodulation, [], [f1337, f2158])).
fof(f1337, plain, (m1_filter_0(k3_filter_0(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | ~ spl28_26), inference(avatar_component_clause, [], [f1335])).
fof(f1335, plain, (spl28_26 <=> m1_filter_0(k3_filter_0(k1_lattice2(sK1), sK2), k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_26])])).
fof(f16142, plain, (! [X25] : (~ m1_filter_0(X25, k1_lattice2(sK1)) | (k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), sK3, X25)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK3), X25)))) | (spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_69 | ~ spl28_80)), inference(backward_demodulation, [], [f8704, f15995])).
fof(f8704, plain, (! [X25] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK3, X25)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k19_filter_2(sK1, sK3), X25))) | ~ m1_filter_0(X25, k1_lattice2(sK1))) | (spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_80)), inference(forward_demodulation, [], [f8703, f2136])).
fof(f8703, plain, (! [X25] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK3, X25)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k3_filter_0(k1_lattice2(sK1), sK3), X25))) | ~ m1_filter_0(X25, k1_lattice2(sK1))) | (spl28_21 | ~ spl28_22 | ~ spl28_23)), inference(subsumption_resolution, [], [f8702, f1312])).
fof(f8702, plain, (! [X25] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK3, X25)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k3_filter_0(k1_lattice2(sK1), sK3), X25))) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_filter_0(X25, k1_lattice2(sK1))) | (~ spl28_22 | ~ spl28_23)), inference(subsumption_resolution, [], [f8701, f1316])).
fof(f8701, plain, (! [X25] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK3, X25)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k3_filter_0(k1_lattice2(sK1), sK3), X25))) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_filter_0(X25, k1_lattice2(sK1))) | ~ spl28_23), inference(subsumption_resolution, [], [f8700, f1320])).
fof(f8700, plain, ! [X25] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK3, X25)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k3_filter_0(k1_lattice2(sK1), sK3), X25))) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_filter_0(X25, k1_lattice2(sK1))), inference(subsumption_resolution, [], [f8665, f300])).
fof(f8665, plain, ! [X25] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), sK3, X25)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), k3_filter_0(k1_lattice2(sK1), sK3), X25))) | v1_xboole_0(sK3) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_filter_0(X25, k1_lattice2(sK1))), inference(resolution, [], [f1384, f1085])).
fof(f427993, plain, (~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | spl28_2), inference(forward_demodulation, [], [f472, f15451])).
fof(f472, plain, (~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, k19_filter_2(sK1, sK3)))) | spl28_2), inference(avatar_component_clause, [], [f470])).
fof(f470, plain, (spl28_2 <=> r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, k19_filter_2(sK1, sK3))))), introduced(avatar_definition, [new_symbols(naming, [spl28_2])])).
fof(f427998, plain, (r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), k19_filter_2(sK1, k2_xboole_0(sK2, sK3))) | (~ spl28_1 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(forward_demodulation, [], [f427997, f2173])).
fof(f427997, plain, (r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k2_xboole_0(sK2, sK3))) | (~ spl28_1 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(forward_demodulation, [], [f427996, f427919])).
fof(f427919, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(backward_demodulation, [], [f365862, f427898])).
fof(f365862, plain, ((k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | (~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_69 | ~ spl28_80 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(backward_demodulation, [], [f222069, f365861])).
fof(f222069, plain, ((k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (~ spl28_2170 | ~ spl28_2183)), inference(backward_demodulation, [], [f221935, f221802])).
fof(f221802, plain, ((k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | ~ spl28_2170), inference(avatar_component_clause, [], [f221800])).
fof(f221800, plain, (spl28_2170 <=> (k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))), introduced(avatar_definition, [new_symbols(naming, [spl28_2170])])).
fof(f221935, plain, ((k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | ~ spl28_2183), inference(avatar_component_clause, [], [f221933])).
fof(f221933, plain, (spl28_2183 <=> (k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))), introduced(avatar_definition, [new_symbols(naming, [spl28_2183])])).
fof(f427996, plain, (r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | ~ spl28_1), inference(forward_demodulation, [], [f467, f15326])).
fof(f15326, plain, (k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3) = k2_xboole_0(sK3, k19_filter_2(sK1, sK2))), inference(forward_demodulation, [], [f15313, f318])).
fof(f15313, plain, (k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3) = k2_xboole_0(k19_filter_2(sK1, sK2), sK3)), inference(resolution, [], [f2181, f954])).
fof(f2181, plain, ! [X0] : (~ m2_lattice4(X0, sK1) | (k2_xboole_0(X0, sK3) = k4_subset_1(u1_struct_0(sK1), X0, sK3))), inference(subsumption_resolution, [], [f2180, f295])).
fof(f2180, plain, ! [X0] : ((k2_xboole_0(X0, sK3) = k4_subset_1(u1_struct_0(sK1), X0, sK3)) | ~ m2_lattice4(X0, sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f2179, f296])).
fof(f2179, plain, ! [X0] : ((k2_xboole_0(X0, sK3) = k4_subset_1(u1_struct_0(sK1), X0, sK3)) | ~ m2_lattice4(X0, sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f2167, f297])).
fof(f2167, plain, ! [X0] : ((k2_xboole_0(X0, sK3) = k4_subset_1(u1_struct_0(sK1), X0, sK3)) | ~ m2_lattice4(X0, sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f836, f344])).
fof(f467, plain, (r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3))) | ~ spl28_1), inference(avatar_component_clause, [], [f466])).
fof(f466, plain, (spl28_1 <=> r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3)))), introduced(avatar_definition, [new_symbols(naming, [spl28_1])])).
fof(f427987, plain, (spl28_1 | spl28_14 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_246 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183), inference(avatar_contradiction_clause, [], [f427986])).
fof(f427986, plain, ($false | (spl28_1 | spl28_14 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_246 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(subsumption_resolution, [], [f427985, f674])).
fof(f674, plain, (~ v1_xboole_0(u1_struct_0(sK1)) | spl28_14), inference(avatar_component_clause, [], [f672])).
fof(f672, plain, (spl28_14 <=> v1_xboole_0(u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_14])])).
fof(f427985, plain, (v1_xboole_0(u1_struct_0(sK1)) | (spl28_1 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_246 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(subsumption_resolution, [], [f427984, f11393])).
fof(f11393, plain, (m2_lattice4(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | ~ spl28_246), inference(subsumption_resolution, [], [f11392, f295])).
fof(f11392, plain, (m2_lattice4(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v3_struct_0(sK1) | ~ spl28_246), inference(subsumption_resolution, [], [f11391, f296])).
fof(f11391, plain, (m2_lattice4(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ spl28_246), inference(subsumption_resolution, [], [f11386, f297])).
fof(f11386, plain, (m2_lattice4(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ spl28_246), inference(resolution, [], [f11120, f343])).
fof(f11120, plain, (m2_filter_2(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | ~ spl28_246), inference(avatar_component_clause, [], [f11118])).
fof(f11118, plain, (spl28_246 <=> m2_filter_2(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1)), introduced(avatar_definition, [new_symbols(naming, [spl28_246])])).
fof(f427984, plain, (~ m2_lattice4(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v1_xboole_0(u1_struct_0(sK1)) | (spl28_1 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(subsumption_resolution, [], [f427983, f295])).
fof(f427983, plain, (v3_struct_0(sK1) | ~ m2_lattice4(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v1_xboole_0(u1_struct_0(sK1)) | (spl28_1 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(subsumption_resolution, [], [f427982, f296])).
fof(f427982, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m2_lattice4(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v1_xboole_0(u1_struct_0(sK1)) | (spl28_1 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(subsumption_resolution, [], [f427975, f297])).
fof(f427975, plain, (~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m2_lattice4(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v1_xboole_0(u1_struct_0(sK1)) | (spl28_1 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(resolution, [], [f427920, f768])).
fof(f768, plain, ! [X0, X1] : (r1_filter_2(u1_struct_0(X1), X0, X0) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ m2_lattice4(X0, X1) | v1_xboole_0(u1_struct_0(X1))), inference(resolution, [], [f344, f464])).
fof(f464, plain, ! [X2, X0] : (~ m1_subset_1(X2, k1_zfmisc_1(X0)) | r1_filter_2(X0, X2, X2) | v1_xboole_0(X0)), inference(duplicate_literal_removal, [], [f463])).
fof(f463, plain, ! [X2, X0] : (r1_filter_2(X0, X2, X2) | ~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X2, k1_zfmisc_1(X0)) | v1_xboole_0(X0)), inference(equality_resolution, [], [f444])).
fof(f444, plain, ! [X2, X0, X1] : (r1_filter_2(X0, X1, X2) | ~ (X1 = X2) | ~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f294])).
fof(f294, plain, ! [X0, X1, X2] : (((r1_filter_2(X0, X1, X2) | ~ (X1 = X2)) & ((X1 = X2) | ~ r1_filter_2(X0, X1, X2))) | ~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0)) | v1_xboole_0(X0)), inference(nnf_transformation, [], [f218])).
fof(f218, plain, ! [X0, X1, X2] : ((r1_filter_2(X0, X1, X2) <=> (X1 = X2)) | ~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0)) | v1_xboole_0(X0)), inference(flattening, [], [f217])).
fof(f217, plain, ! [X0, X1, X2] : ((r1_filter_2(X0, X1, X2) <=> (X1 = X2)) | (~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0)) | v1_xboole_0(X0))), inference(ennf_transformation, [], [f89])).
fof(f89, plain, ! [X0, X1, X2] : ((m1_subset_1(X2, k1_zfmisc_1(X0)) & m1_subset_1(X1, k1_zfmisc_1(X0)) & ~ v1_xboole_0(X0)) => (r1_filter_2(X0, X1, X2) <=> (X1 = X2))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', redefinition_r1_filter_2)).
fof(f427920, plain, (~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), k19_filter_2(sK1, k2_xboole_0(sK2, sK3))) | (spl28_1 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_32 | ~ spl28_69 | ~ spl28_80 | ~ spl28_285 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(backward_demodulation, [], [f365863, f427898])).
fof(f365863, plain, (~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | (spl28_1 | ~ spl28_19 | spl28_21 | ~ spl28_22 | ~ spl28_23 | ~ spl28_24 | ~ spl28_26 | ~ spl28_69 | ~ spl28_80 | ~ spl28_2141 | ~ spl28_2154 | ~ spl28_2170 | ~ spl28_2183)), inference(backward_demodulation, [], [f15327, f365862])).
fof(f15327, plain, (~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | spl28_1), inference(backward_demodulation, [], [f2191, f15326])).
fof(f2191, plain, (~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3))) | spl28_1), inference(backward_demodulation, [], [f468, f2173])).
fof(f468, plain, (~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3))) | spl28_1), inference(avatar_component_clause, [], [f466])).
fof(f222035, plain, ~ spl28_2162, inference(avatar_contradiction_clause, [], [f222034])).
fof(f222034, plain, ($false | ~ spl28_2162), inference(subsumption_resolution, [], [f222022, f300])).
fof(f222022, plain, (v1_xboole_0(sK3) | ~ spl28_2162), inference(resolution, [], [f221724, f373])).
fof(f373, plain, ! [X0, X1] : (~ v1_xboole_0(k2_xboole_0(X0, X1)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ! [X0, X1] : (~ v1_xboole_0(k2_xboole_0(X0, X1)) | v1_xboole_0(X0)), inference(ennf_transformation, [], [f60])).
fof(f60, plain, ! [X0, X1] : (~ v1_xboole_0(X0) => ~ v1_xboole_0(k2_xboole_0(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', fc2_xboole_0)).
fof(f221724, plain, (v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) | ~ spl28_2162), inference(avatar_component_clause, [], [f221722])).
fof(f221722, plain, (spl28_2162 <=> v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))), introduced(avatar_definition, [new_symbols(naming, [spl28_2162])])).
fof(f221936, plain, (spl28_2162 | spl28_2183 | ~ spl28_22 | ~ spl28_1317), inference(avatar_split_clause, [], [f221931, f161835, f1315, f221933, f221722])).
fof(f161835, plain, (spl28_1317 <=> m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1)))), introduced(avatar_definition, [new_symbols(naming, [spl28_1317])])).
fof(f221931, plain, ((k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) | (~ spl28_22 | ~ spl28_1317)), inference(forward_demodulation, [], [f221930, f163181])).
fof(f163181, plain, ((k2_xboole_0(sK3, k19_filter_2(sK1, sK2)) = k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | ~ spl28_1317), inference(forward_demodulation, [], [f162942, f15326])).
fof(f162942, plain, ((k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3) = k7_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3))) | ~ spl28_1317), inference(resolution, [], [f161836, f6977])).
fof(f6977, plain, ! [X24] : (~ m1_subset_1(X24, k1_zfmisc_1(u1_struct_0(sK1))) | (k4_subset_1(u1_struct_0(sK1), X24, sK3) = k7_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X24, sK3)))), inference(subsumption_resolution, [], [f6976, f295])).
fof(f6976, plain, ! [X24] : (v3_struct_0(sK1) | (k4_subset_1(u1_struct_0(sK1), X24, sK3) = k7_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X24, sK3))) | ~ m1_subset_1(X24, k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f6975, f296])).
fof(f6975, plain, ! [X24] : (~ v10_lattices(sK1) | v3_struct_0(sK1) | (k4_subset_1(u1_struct_0(sK1), X24, sK3) = k7_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X24, sK3))) | ~ m1_subset_1(X24, k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f6949, f297])).
fof(f6949, plain, ! [X24] : (~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (k4_subset_1(u1_struct_0(sK1), X24, sK3) = k7_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X24, sK3))) | ~ m1_subset_1(X24, k1_zfmisc_1(u1_struct_0(sK1)))), inference(resolution, [], [f811, f301])).
fof(f811, plain, ! [X6, X8, X7] : (~ m1_subset_1(X8, k1_zfmisc_1(u1_struct_0(X6))) | ~ l3_lattices(X6) | ~ v10_lattices(X6) | v3_struct_0(X6) | (k4_subset_1(u1_struct_0(X6), X7, X8) = k7_filter_2(X6, k4_subset_1(u1_struct_0(X6), X7, X8))) | ~ m1_subset_1(X7, k1_zfmisc_1(u1_struct_0(X6)))), inference(resolution, [], [f321, f331])).
fof(f331, plain, ! [X2, X0, X1] : (m1_subset_1(k4_subset_1(X0, X1, X2), k1_zfmisc_1(X0)) | ~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0))), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ! [X0, X1, X2] : (m1_subset_1(k4_subset_1(X0, X1, X2), k1_zfmisc_1(X0)) | ~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0))), inference(flattening, [], [f148])).
fof(f148, plain, ! [X0, X1, X2] : (m1_subset_1(k4_subset_1(X0, X1, X2), k1_zfmisc_1(X0)) | (~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0)))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ! [X0, X1, X2] : ((m1_subset_1(X2, k1_zfmisc_1(X0)) & m1_subset_1(X1, k1_zfmisc_1(X0))) => m1_subset_1(k4_subset_1(X0, X1, X2), k1_zfmisc_1(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_k4_subset_1)).
fof(f161836, plain, (m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_1317), inference(avatar_component_clause, [], [f161835])).
fof(f221930, plain, (v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | (~ spl28_22 | ~ spl28_1317)), inference(forward_demodulation, [], [f221929, f163181])).
fof(f221929, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | (~ spl28_22 | ~ spl28_1317)), inference(subsumption_resolution, [], [f221928, f161836])).
fof(f221928, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_22), inference(subsumption_resolution, [], [f221927, f301])).
fof(f221927, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_22), inference(subsumption_resolution, [], [f221926, f295])).
fof(f221926, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | v3_struct_0(sK1) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_22), inference(subsumption_resolution, [], [f221925, f296])).
fof(f221925, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_22), inference(subsumption_resolution, [], [f221924, f297])).
fof(f221924, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_22), inference(subsumption_resolution, [], [f221695, f1316])).
fof(f221695, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | ~ v10_lattices(k1_lattice2(sK1)) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1)))), inference(superposition, [], [f7773, f15326])).
fof(f7773, plain, ! [X6, X8, X7] : (v1_xboole_0(k7_filter_2(X6, k4_subset_1(u1_struct_0(X6), X7, X8))) | ~ v10_lattices(k1_lattice2(X6)) | (k3_filter_2(k1_lattice2(X6), k7_filter_2(X6, k4_subset_1(u1_struct_0(X6), X7, X8))) = k3_filter_0(k1_lattice2(X6), k7_filter_2(X6, k4_subset_1(u1_struct_0(X6), X7, X8)))) | ~ l3_lattices(X6) | ~ v10_lattices(X6) | v3_struct_0(X6) | ~ m1_subset_1(X8, k1_zfmisc_1(u1_struct_0(X6))) | ~ m1_subset_1(X7, k1_zfmisc_1(u1_struct_0(X6)))), inference(resolution, [], [f1136, f331])).
fof(f1136, plain, ! [X10, X9] : (~ m1_subset_1(X10, k1_zfmisc_1(u1_struct_0(X9))) | v1_xboole_0(k7_filter_2(X9, X10)) | ~ v10_lattices(k1_lattice2(X9)) | (k3_filter_0(k1_lattice2(X9), k7_filter_2(X9, X10)) = k3_filter_2(k1_lattice2(X9), k7_filter_2(X9, X10))) | ~ l3_lattices(X9) | ~ v10_lattices(X9) | v3_struct_0(X9)), inference(subsumption_resolution, [], [f1135, f364])).
fof(f364, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ! [X0] : ((v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f178])).
fof(f178, plain, ! [X0] : ((v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : ((l3_lattices(X0) & ~ v3_struct_0(X0)) => (v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', fc1_lattice2)).
fof(f1135, plain, ! [X10, X9] : ((k3_filter_0(k1_lattice2(X9), k7_filter_2(X9, X10)) = k3_filter_2(k1_lattice2(X9), k7_filter_2(X9, X10))) | v1_xboole_0(k7_filter_2(X9, X10)) | ~ v10_lattices(k1_lattice2(X9)) | v3_struct_0(k1_lattice2(X9)) | ~ m1_subset_1(X10, k1_zfmisc_1(u1_struct_0(X9))) | ~ l3_lattices(X9) | ~ v10_lattices(X9) | v3_struct_0(X9)), inference(subsumption_resolution, [], [f1122, f328])).
fof(f1122, plain, ! [X10, X9] : ((k3_filter_0(k1_lattice2(X9), k7_filter_2(X9, X10)) = k3_filter_2(k1_lattice2(X9), k7_filter_2(X9, X10))) | v1_xboole_0(k7_filter_2(X9, X10)) | ~ l3_lattices(k1_lattice2(X9)) | ~ v10_lattices(k1_lattice2(X9)) | v3_struct_0(k1_lattice2(X9)) | ~ m1_subset_1(X10, k1_zfmisc_1(u1_struct_0(X9))) | ~ l3_lattices(X9) | ~ v10_lattices(X9) | v3_struct_0(X9)), inference(resolution, [], [f437, f332])).
fof(f437, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | (k3_filter_0(X0, X1) = k3_filter_2(X0, X1)) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f212])).
fof(f212, plain, ! [X0, X1] : ((k3_filter_0(X0, X1) = k3_filter_2(X0, X1)) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f211])).
fof(f211, plain, ! [X0, X1] : ((k3_filter_0(X0, X1) = k3_filter_2(X0, X1)) | (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f85])).
fof(f85, plain, ! [X0, X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k3_filter_0(X0, X1) = k3_filter_2(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', redefinition_k3_filter_2)).
fof(f221803, plain, (spl28_2162 | spl28_2170 | ~ spl28_1317), inference(avatar_split_clause, [], [f221798, f161835, f221800, f221722])).
fof(f221798, plain, ((k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK3, k19_filter_2(sK1, sK2)))) | v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) | ~ spl28_1317), inference(forward_demodulation, [], [f221797, f163181])).
fof(f221797, plain, (v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) | (k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | ~ spl28_1317), inference(subsumption_resolution, [], [f221796, f161836])).
fof(f221796, plain, (v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) | (k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f221795, f301])).
fof(f221795, plain, (v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) | (k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f221794, f295])).
fof(f221794, plain, (v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) | (k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | v3_struct_0(sK1) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f221793, f296])).
fof(f221793, plain, (v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) | (k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f221680, f297])).
fof(f221680, plain, (v1_xboole_0(k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) | (k19_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK3, k19_filter_2(sK1, sK2))))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1)))), inference(superposition, [], [f5353, f15326])).
fof(f5353, plain, ! [X6, X8, X7] : (v1_xboole_0(k4_subset_1(u1_struct_0(X6), X7, X8)) | (k19_filter_2(X6, k4_subset_1(u1_struct_0(X6), X7, X8)) = k3_filter_2(k1_lattice2(X6), k7_filter_2(X6, k4_subset_1(u1_struct_0(X6), X7, X8)))) | ~ l3_lattices(X6) | ~ v10_lattices(X6) | v3_struct_0(X6) | ~ m1_subset_1(X8, k1_zfmisc_1(u1_struct_0(X6))) | ~ m1_subset_1(X7, k1_zfmisc_1(u1_struct_0(X6)))), inference(resolution, [], [f1170, f331])).
fof(f1170, plain, ! [X26, X25] : (~ m1_subset_1(X26, k1_zfmisc_1(u1_struct_0(X25))) | (k19_filter_2(X25, X26) = k3_filter_2(k1_lattice2(X25), k7_filter_2(X25, X26))) | v1_xboole_0(X26) | ~ l3_lattices(X25) | ~ v10_lattices(X25) | v3_struct_0(X25)), inference(subsumption_resolution, [], [f1169, f364])).
fof(f1169, plain, ! [X26, X25] : ((k19_filter_2(X25, X26) = k3_filter_2(k1_lattice2(X25), k7_filter_2(X25, X26))) | ~ m1_subset_1(X26, k1_zfmisc_1(u1_struct_0(X25))) | v1_xboole_0(X26) | ~ l3_lattices(X25) | ~ v10_lattices(X25) | v3_struct_0(X25) | v3_struct_0(k1_lattice2(X25))), inference(subsumption_resolution, [], [f1168, f514])).
fof(f514, plain, ! [X0] : (l1_struct_0(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(resolution, [], [f495, f335])).
fof(f335, plain, ! [X0] : (~ l2_lattices(X0) | l1_struct_0(X0)), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ! [X0] : (l1_struct_0(X0) | ~ l2_lattices(X0)), inference(ennf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : (l2_lattices(X0) => l1_struct_0(X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_l2_lattices)).
fof(f495, plain, ! [X0] : (l2_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(resolution, [], [f328, f337])).
fof(f337, plain, ! [X0] : (~ l3_lattices(X0) | l2_lattices(X0)), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ! [X0] : ((l2_lattices(X0) & l1_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_l3_lattices)).
fof(f1168, plain, ! [X26, X25] : ((k19_filter_2(X25, X26) = k3_filter_2(k1_lattice2(X25), k7_filter_2(X25, X26))) | ~ m1_subset_1(X26, k1_zfmisc_1(u1_struct_0(X25))) | v1_xboole_0(X26) | ~ l3_lattices(X25) | ~ v10_lattices(X25) | v3_struct_0(X25) | ~ l1_struct_0(k1_lattice2(X25)) | v3_struct_0(k1_lattice2(X25))), inference(subsumption_resolution, [], [f1158, f422])).
fof(f422, plain, ! [X0] : (~ v1_xboole_0(sK25(X0)) | ~ l1_struct_0(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f287])).
fof(f287, plain, ! [X0] : ((~ v1_xboole_0(sK25(X0)) & m1_subset_1(sK25(X0), k1_zfmisc_1(u1_struct_0(X0)))) | ~ l1_struct_0(X0) | v3_struct_0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK25])], [f208, f286])).
fof(f286, plain, ! [X0] : (? [X1] : (~ v1_xboole_0(X1) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) => (~ v1_xboole_0(sK25(X0)) & m1_subset_1(sK25(X0), k1_zfmisc_1(u1_struct_0(X0))))), introduced(choice_axiom, [])).
fof(f208, plain, ! [X0] : (? [X1] : (~ v1_xboole_0(X1) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | ~ l1_struct_0(X0) | v3_struct_0(X0)), inference(flattening, [], [f207])).
fof(f207, plain, ! [X0] : (? [X1] : (~ v1_xboole_0(X1) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | (~ l1_struct_0(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f81])).
fof(f81, plain, ! [X0] : ((l1_struct_0(X0) & ~ v3_struct_0(X0)) => ? [X1] : (~ v1_xboole_0(X1) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', rc5_struct_0)).
fof(f1158, plain, ! [X26, X25] : ((k19_filter_2(X25, X26) = k3_filter_2(k1_lattice2(X25), k7_filter_2(X25, X26))) | v1_xboole_0(sK25(k1_lattice2(X25))) | ~ m1_subset_1(X26, k1_zfmisc_1(u1_struct_0(X25))) | v1_xboole_0(X26) | ~ l3_lattices(X25) | ~ v10_lattices(X25) | v3_struct_0(X25) | ~ l1_struct_0(k1_lattice2(X25)) | v3_struct_0(k1_lattice2(X25))), inference(resolution, [], [f451, f421])).
fof(f421, plain, ! [X0] : (m1_subset_1(sK25(X0), k1_zfmisc_1(u1_struct_0(X0))) | ~ l1_struct_0(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f287])).
fof(f451, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) | (k19_filter_2(X0, X1) = k3_filter_2(k1_lattice2(X0), k7_filter_2(X0, X1))) | v1_xboole_0(X2) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f227])).
fof(f227, plain, ! [X0] : (! [X1] : (! [X2] : (((k3_filter_2(k1_lattice2(X0), X2) = k19_filter_2(X0, k8_filter_2(X0, X2))) & (k3_filter_2(X0, k8_filter_2(X0, X2)) = k19_filter_2(k1_lattice2(X0), X2)) & (k3_filter_2(X0, X1) = k19_filter_2(k1_lattice2(X0), k7_filter_2(X0, X1))) & (k19_filter_2(X0, X1) = k3_filter_2(k1_lattice2(X0), k7_filter_2(X0, X1)))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) | v1_xboole_0(X2)) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f226])).
fof(f226, plain, ! [X0] : (! [X1] : (! [X2] : (((k3_filter_2(k1_lattice2(X0), X2) = k19_filter_2(X0, k8_filter_2(X0, X2))) & (k3_filter_2(X0, k8_filter_2(X0, X2)) = k19_filter_2(k1_lattice2(X0), X2)) & (k3_filter_2(X0, X1) = k19_filter_2(k1_lattice2(X0), k7_filter_2(X0, X1))) & (k19_filter_2(X0, X1) = k3_filter_2(k1_lattice2(X0), k7_filter_2(X0, X1)))) | (~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) | v1_xboole_0(X2))) | (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f96])).
fof(f96, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) => ! [X2] : ((m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) & ~ v1_xboole_0(X2)) => ((k3_filter_2(k1_lattice2(X0), X2) = k19_filter_2(X0, k8_filter_2(X0, X2))) & (k3_filter_2(X0, k8_filter_2(X0, X2)) = k19_filter_2(k1_lattice2(X0), X2)) & (k3_filter_2(X0, X1) = k19_filter_2(k1_lattice2(X0), k7_filter_2(X0, X1))) & (k19_filter_2(X0, X1) = k3_filter_2(k1_lattice2(X0), k7_filter_2(X0, X1))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', t37_filter_2)).
fof(f221505, plain, ~ spl28_2133, inference(avatar_contradiction_clause, [], [f221504])).
fof(f221504, plain, ($false | ~ spl28_2133), inference(subsumption_resolution, [], [f221492, f298])).
fof(f221492, plain, (v1_xboole_0(sK2) | ~ spl28_2133), inference(resolution, [], [f221194, f373])).
fof(f221194, plain, (v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) | ~ spl28_2133), inference(avatar_component_clause, [], [f221192])).
fof(f221192, plain, (spl28_2133 <=> v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))), introduced(avatar_definition, [new_symbols(naming, [spl28_2133])])).
fof(f221406, plain, (spl28_2133 | spl28_2154 | ~ spl28_22 | ~ spl28_1328), inference(avatar_split_clause, [], [f221401, f169153, f1315, f221403, f221192])).
fof(f169153, plain, (spl28_1328 <=> m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1)))), introduced(avatar_definition, [new_symbols(naming, [spl28_1328])])).
fof(f221401, plain, ((k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) | (~ spl28_22 | ~ spl28_1328)), inference(forward_demodulation, [], [f221400, f170717])).
fof(f170717, plain, ((k2_xboole_0(sK2, k19_filter_2(sK1, sK3)) = k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | ~ spl28_1328), inference(forward_demodulation, [], [f170465, f15215])).
fof(f15215, plain, (k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK3), sK2) = k2_xboole_0(sK2, k19_filter_2(sK1, sK3))), inference(forward_demodulation, [], [f15202, f318])).
fof(f15202, plain, (k2_xboole_0(k19_filter_2(sK1, sK3), sK2) = k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK3), sK2)), inference(resolution, [], [f2067, f1014])).
fof(f2067, plain, ! [X0] : (~ m2_lattice4(X0, sK1) | (k2_xboole_0(X0, sK2) = k4_subset_1(u1_struct_0(sK1), X0, sK2))), inference(subsumption_resolution, [], [f2066, f295])).
fof(f2066, plain, ! [X0] : ((k2_xboole_0(X0, sK2) = k4_subset_1(u1_struct_0(sK1), X0, sK2)) | ~ m2_lattice4(X0, sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f2065, f296])).
fof(f2065, plain, ! [X0] : ((k2_xboole_0(X0, sK2) = k4_subset_1(u1_struct_0(sK1), X0, sK2)) | ~ m2_lattice4(X0, sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f2053, f297])).
fof(f2053, plain, ! [X0] : ((k2_xboole_0(X0, sK2) = k4_subset_1(u1_struct_0(sK1), X0, sK2)) | ~ m2_lattice4(X0, sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f835, f344])).
fof(f170465, plain, ((k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK3), sK2) = k7_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK3), sK2))) | ~ spl28_1328), inference(resolution, [], [f169154, f6968])).
fof(f6968, plain, ! [X21] : (~ m1_subset_1(X21, k1_zfmisc_1(u1_struct_0(sK1))) | (k4_subset_1(u1_struct_0(sK1), X21, sK2) = k7_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X21, sK2)))), inference(subsumption_resolution, [], [f6967, f295])).
fof(f6967, plain, ! [X21] : (v3_struct_0(sK1) | (k4_subset_1(u1_struct_0(sK1), X21, sK2) = k7_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X21, sK2))) | ~ m1_subset_1(X21, k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f6966, f296])).
fof(f6966, plain, ! [X21] : (~ v10_lattices(sK1) | v3_struct_0(sK1) | (k4_subset_1(u1_struct_0(sK1), X21, sK2) = k7_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X21, sK2))) | ~ m1_subset_1(X21, k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f6946, f297])).
fof(f6946, plain, ! [X21] : (~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (k4_subset_1(u1_struct_0(sK1), X21, sK2) = k7_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), X21, sK2))) | ~ m1_subset_1(X21, k1_zfmisc_1(u1_struct_0(sK1)))), inference(resolution, [], [f811, f299])).
fof(f169154, plain, (m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_1328), inference(avatar_component_clause, [], [f169153])).
fof(f221400, plain, (v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | (~ spl28_22 | ~ spl28_1328)), inference(forward_demodulation, [], [f221399, f170717])).
fof(f221399, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | (~ spl28_22 | ~ spl28_1328)), inference(subsumption_resolution, [], [f221398, f169154])).
fof(f221398, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_22), inference(subsumption_resolution, [], [f221397, f299])).
fof(f221397, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_22), inference(subsumption_resolution, [], [f221396, f295])).
fof(f221396, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | v3_struct_0(sK1) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_22), inference(subsumption_resolution, [], [f221395, f296])).
fof(f221395, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_22), inference(subsumption_resolution, [], [f221394, f297])).
fof(f221394, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1))) | ~ spl28_22), inference(subsumption_resolution, [], [f221165, f1316])).
fof(f221165, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | ~ v10_lattices(k1_lattice2(sK1)) | (k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) = k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1)))), inference(superposition, [], [f7773, f15215])).
fof(f221273, plain, (spl28_2133 | spl28_2141 | ~ spl28_1328), inference(avatar_split_clause, [], [f221268, f169153, f221270, f221192])).
fof(f221268, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK2, k19_filter_2(sK1, sK3)))) | v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) | ~ spl28_1328), inference(forward_demodulation, [], [f221267, f170717])).
fof(f221267, plain, (v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) | (k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | ~ spl28_1328), inference(subsumption_resolution, [], [f221266, f169154])).
fof(f221266, plain, (v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) | (k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f221265, f299])).
fof(f221265, plain, (v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) | (k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f221264, f295])).
fof(f221264, plain, (v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) | (k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | v3_struct_0(sK1) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f221263, f296])).
fof(f221263, plain, (v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) | (k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f221150, f297])).
fof(f221150, plain, (v1_xboole_0(k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) | (k19_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, k19_filter_2(sK1, sK3))))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1)))), inference(superposition, [], [f5353, f15215])).
fof(f169377, plain, spl28_1328, inference(avatar_contradiction_clause, [], [f169376])).
fof(f169376, plain, ($false | spl28_1328), inference(subsumption_resolution, [], [f169375, f295])).
fof(f169375, plain, (v3_struct_0(sK1) | spl28_1328), inference(subsumption_resolution, [], [f169374, f296])).
fof(f169374, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | spl28_1328), inference(subsumption_resolution, [], [f169373, f297])).
fof(f169373, plain, (~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | spl28_1328), inference(subsumption_resolution, [], [f169370, f1014])).
fof(f169370, plain, (~ m2_lattice4(k19_filter_2(sK1, sK3), sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | spl28_1328), inference(resolution, [], [f169155, f344])).
fof(f169155, plain, (~ m1_subset_1(k19_filter_2(sK1, sK3), k1_zfmisc_1(u1_struct_0(sK1))) | spl28_1328), inference(avatar_component_clause, [], [f169153])).
fof(f162084, plain, spl28_1317, inference(avatar_contradiction_clause, [], [f162083])).
fof(f162083, plain, ($false | spl28_1317), inference(subsumption_resolution, [], [f162082, f295])).
fof(f162082, plain, (v3_struct_0(sK1) | spl28_1317), inference(subsumption_resolution, [], [f162081, f296])).
fof(f162081, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | spl28_1317), inference(subsumption_resolution, [], [f162080, f297])).
fof(f162080, plain, (~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | spl28_1317), inference(subsumption_resolution, [], [f162077, f954])).
fof(f162077, plain, (~ m2_lattice4(k19_filter_2(sK1, sK2), sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | spl28_1317), inference(resolution, [], [f161837, f344])).
fof(f161837, plain, (~ m1_subset_1(k19_filter_2(sK1, sK2), k1_zfmisc_1(u1_struct_0(sK1))) | spl28_1317), inference(avatar_component_clause, [], [f161835])).
fof(f13029, plain, spl28_66, inference(avatar_contradiction_clause, [], [f13028])).
fof(f13028, plain, ($false | spl28_66), inference(subsumption_resolution, [], [f13027, f486])).
fof(f486, plain, l2_lattices(sK1), inference(resolution, [], [f337, f297])).
fof(f13027, plain, (~ l2_lattices(sK1) | spl28_66), inference(resolution, [], [f1644, f739])).
fof(f739, plain, ! [X1] : (m1_relset_1(u2_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l2_lattices(X1)), inference(resolution, [], [f351, f441])).
fof(f441, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f293])).
fof(f293, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f88])).
fof(f88, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', redefinition_m2_relset_1)).
fof(f351, plain, ! [X0] : (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ! [X0] : ((m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0))) | ~ l2_lattices(X0)), inference(ennf_transformation, [], [f41])).
fof(f41, plain, ! [X0] : (l2_lattices(X0) => (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_u2_lattices)).
fof(f1644, plain, (~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl28_66), inference(avatar_component_clause, [], [f1642])).
fof(f1642, plain, (spl28_66 <=> m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_66])])).
fof(f13010, plain, spl28_65, inference(avatar_contradiction_clause, [], [f13009])).
fof(f13009, plain, ($false | spl28_65), inference(subsumption_resolution, [], [f13008, f486])).
fof(f13008, plain, (~ l2_lattices(sK1) | spl28_65), inference(resolution, [], [f1640, f350])).
fof(f350, plain, ! [X0] : (v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f167])).
fof(f1640, plain, (~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl28_65), inference(avatar_component_clause, [], [f1638])).
fof(f1638, plain, (spl28_65 <=> v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_65])])).
fof(f12169, plain, spl28_64, inference(avatar_contradiction_clause, [], [f12168])).
fof(f12168, plain, ($false | spl28_64), inference(subsumption_resolution, [], [f12167, f486])).
fof(f12167, plain, (~ l2_lattices(sK1) | spl28_64), inference(resolution, [], [f1636, f349])).
fof(f349, plain, ! [X0] : (v1_funct_1(u2_lattices(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f167])).
fof(f1636, plain, (~ v1_funct_1(u2_lattices(sK1)) | spl28_64), inference(avatar_component_clause, [], [f1634])).
fof(f1634, plain, (spl28_64 <=> v1_funct_1(u2_lattices(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_64])])).
fof(f12166, plain, spl28_63, inference(avatar_contradiction_clause, [], [f12165])).
fof(f12165, plain, ($false | spl28_63), inference(subsumption_resolution, [], [f12164, f480])).
fof(f480, plain, l1_lattices(sK1), inference(resolution, [], [f336, f297])).
fof(f336, plain, ! [X0] : (~ l3_lattices(X0) | l1_lattices(X0)), inference(cnf_transformation, [], [f156])).
fof(f12164, plain, (~ l1_lattices(sK1) | spl28_63), inference(resolution, [], [f1632, f737])).
fof(f737, plain, ! [X1] : (m1_relset_1(u1_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l1_lattices(X1)), inference(resolution, [], [f348, f441])).
fof(f348, plain, ! [X0] : (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ! [X0] : ((m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0))) | ~ l1_lattices(X0)), inference(ennf_transformation, [], [f39])).
fof(f39, plain, ! [X0] : (l1_lattices(X0) => (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', dt_u1_lattices)).
fof(f1632, plain, (~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl28_63), inference(avatar_component_clause, [], [f1630])).
fof(f1630, plain, (spl28_63 <=> m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_63])])).
fof(f12151, plain, spl28_62, inference(avatar_contradiction_clause, [], [f12150])).
fof(f12150, plain, ($false | spl28_62), inference(subsumption_resolution, [], [f12149, f480])).
fof(f12149, plain, (~ l1_lattices(sK1) | spl28_62), inference(resolution, [], [f1628, f347])).
fof(f347, plain, ! [X0] : (v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f166])).
fof(f1628, plain, (~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl28_62), inference(avatar_component_clause, [], [f1626])).
fof(f1626, plain, (spl28_62 <=> v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_62])])).
fof(f11816, plain, (spl28_285 | ~ spl28_22 | spl28_30 | spl28_243), inference(avatar_split_clause, [], [f11815, f11088, f1357, f1315, f11794])).
fof(f1357, plain, (spl28_30 <=> v1_xboole_0(u1_struct_0(k1_lattice2(sK1)))), introduced(avatar_definition, [new_symbols(naming, [spl28_30])])).
fof(f11088, plain, (spl28_243 <=> v1_xboole_0(k2_xboole_0(sK2, sK3))), introduced(avatar_definition, [new_symbols(naming, [spl28_243])])).
fof(f11815, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, sK3))) | (~ spl28_22 | spl28_30 | spl28_243)), inference(backward_demodulation, [], [f11777, f11814])).
fof(f11814, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK2, sK3))) | (spl28_30 | spl28_243)), inference(forward_demodulation, [], [f11813, f11703])).
fof(f11703, plain, (k2_xboole_0(sK2, sK3) = k7_filter_2(sK1, k2_xboole_0(sK2, sK3))), inference(subsumption_resolution, [], [f11702, f295])).
fof(f11702, plain, ((k2_xboole_0(sK2, sK3) = k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f11701, f296])).
fof(f11701, plain, ((k2_xboole_0(sK2, sK3) = k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f11612, f297])).
fof(f11612, plain, ((k2_xboole_0(sK2, sK3) = k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f11140, f321])).
fof(f11140, plain, m1_subset_1(k2_xboole_0(sK2, sK3), k1_zfmisc_1(u1_struct_0(sK1))), inference(subsumption_resolution, [], [f11139, f301])).
fof(f11139, plain, (m1_subset_1(k2_xboole_0(sK2, sK3), k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f11077, f299])).
fof(f11077, plain, (m1_subset_1(k2_xboole_0(sK2, sK3), k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1)))), inference(superposition, [], [f331, f2078])).
fof(f11813, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | (spl28_30 | spl28_243)), inference(subsumption_resolution, [], [f11812, f1358])).
fof(f1358, plain, (~ v1_xboole_0(u1_struct_0(k1_lattice2(sK1))) | spl28_30), inference(avatar_component_clause, [], [f1357])).
fof(f11812, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | v1_xboole_0(u1_struct_0(k1_lattice2(sK1))) | spl28_243), inference(subsumption_resolution, [], [f11811, f295])).
fof(f11811, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | v3_struct_0(sK1) | v1_xboole_0(u1_struct_0(k1_lattice2(sK1))) | spl28_243), inference(subsumption_resolution, [], [f11810, f296])).
fof(f11810, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | v1_xboole_0(u1_struct_0(k1_lattice2(sK1))) | spl28_243), inference(subsumption_resolution, [], [f11809, f297])).
fof(f11809, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | v1_xboole_0(u1_struct_0(k1_lattice2(sK1))) | spl28_243), inference(subsumption_resolution, [], [f11636, f11089])).
fof(f11089, plain, (~ v1_xboole_0(k2_xboole_0(sK2, sK3)) | spl28_243), inference(avatar_component_clause, [], [f11088])).
fof(f11636, plain, ((k19_filter_2(sK1, k2_xboole_0(sK2, sK3)) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | v1_xboole_0(k2_xboole_0(sK2, sK3)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | v1_xboole_0(u1_struct_0(k1_lattice2(sK1)))), inference(resolution, [], [f11140, f1167])).
fof(f1167, plain, ! [X24, X23] : (~ m1_subset_1(X24, k1_zfmisc_1(u1_struct_0(X23))) | (k19_filter_2(X23, X24) = k3_filter_2(k1_lattice2(X23), k7_filter_2(X23, X24))) | v1_xboole_0(X24) | ~ l3_lattices(X23) | ~ v10_lattices(X23) | v3_struct_0(X23) | v1_xboole_0(u1_struct_0(k1_lattice2(X23)))), inference(subsumption_resolution, [], [f1157, f407])).
fof(f407, plain, ! [X0] : (~ v1_xboole_0(sK17(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ! [X0] : ((~ v1_xboole_0(sK17(X0)) & m1_subset_1(sK17(X0), k1_zfmisc_1(X0))) | v1_xboole_0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK17])], [f206, f270])).
fof(f270, plain, ! [X0] : (? [X1] : (~ v1_xboole_0(X1) & m1_subset_1(X1, k1_zfmisc_1(X0))) => (~ v1_xboole_0(sK17(X0)) & m1_subset_1(sK17(X0), k1_zfmisc_1(X0)))), introduced(choice_axiom, [])).
fof(f206, plain, ! [X0] : (? [X1] : (~ v1_xboole_0(X1) & m1_subset_1(X1, k1_zfmisc_1(X0))) | v1_xboole_0(X0)), inference(ennf_transformation, [], [f73])).
fof(f73, plain, ! [X0] : (~ v1_xboole_0(X0) => ? [X1] : (~ v1_xboole_0(X1) & m1_subset_1(X1, k1_zfmisc_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', rc1_subset_1)).
fof(f1157, plain, ! [X24, X23] : ((k19_filter_2(X23, X24) = k3_filter_2(k1_lattice2(X23), k7_filter_2(X23, X24))) | v1_xboole_0(sK17(u1_struct_0(k1_lattice2(X23)))) | ~ m1_subset_1(X24, k1_zfmisc_1(u1_struct_0(X23))) | v1_xboole_0(X24) | ~ l3_lattices(X23) | ~ v10_lattices(X23) | v3_struct_0(X23) | v1_xboole_0(u1_struct_0(k1_lattice2(X23)))), inference(resolution, [], [f451, f406])).
fof(f406, plain, ! [X0] : (m1_subset_1(sK17(X0), k1_zfmisc_1(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f271])).
fof(f11777, plain, ((k3_filter_0(k1_lattice2(sK1), k2_xboole_0(sK2, sK3)) = k3_filter_2(k1_lattice2(sK1), k2_xboole_0(sK2, sK3))) | (~ spl28_22 | spl28_243)), inference(forward_demodulation, [], [f11776, f11703])).
fof(f11776, plain, ((k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | (~ spl28_22 | spl28_243)), inference(subsumption_resolution, [], [f11775, f11089])).
fof(f11775, plain, (v1_xboole_0(k2_xboole_0(sK2, sK3)) | (k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | ~ spl28_22), inference(forward_demodulation, [], [f11774, f11703])).
fof(f11774, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) | (k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | ~ spl28_22), inference(subsumption_resolution, [], [f11773, f295])).
fof(f11773, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) | (k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | v3_struct_0(sK1) | ~ spl28_22), inference(subsumption_resolution, [], [f11772, f296])).
fof(f11772, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) | (k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ spl28_22), inference(subsumption_resolution, [], [f11771, f297])).
fof(f11771, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) | (k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ spl28_22), inference(subsumption_resolution, [], [f11631, f1316])).
fof(f11631, plain, (v1_xboole_0(k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) | ~ v10_lattices(k1_lattice2(sK1)) | (k3_filter_0(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3))) = k3_filter_2(k1_lattice2(sK1), k7_filter_2(sK1, k2_xboole_0(sK2, sK3)))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f11140, f1136])).
fof(f11150, plain, ~ spl28_243, inference(avatar_contradiction_clause, [], [f11149])).
fof(f11149, plain, ($false | ~ spl28_243), inference(subsumption_resolution, [], [f11142, f298])).
fof(f11142, plain, (v1_xboole_0(sK2) | ~ spl28_243), inference(resolution, [], [f11090, f373])).
fof(f11090, plain, (v1_xboole_0(k2_xboole_0(sK2, sK3)) | ~ spl28_243), inference(avatar_component_clause, [], [f11088])).
fof(f11121, plain, (spl28_243 | spl28_246), inference(avatar_split_clause, [], [f11116, f11118, f11088])).
fof(f11116, plain, (m2_filter_2(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v1_xboole_0(k2_xboole_0(sK2, sK3))), inference(subsumption_resolution, [], [f11115, f301])).
fof(f11115, plain, (m2_filter_2(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v1_xboole_0(k2_xboole_0(sK2, sK3)) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f11114, f299])).
fof(f11114, plain, (m2_filter_2(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v1_xboole_0(k2_xboole_0(sK2, sK3)) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f11113, f295])).
fof(f11113, plain, (m2_filter_2(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v1_xboole_0(k2_xboole_0(sK2, sK3)) | v3_struct_0(sK1) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f11112, f296])).
fof(f11112, plain, (m2_filter_2(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v1_xboole_0(k2_xboole_0(sK2, sK3)) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1)))), inference(subsumption_resolution, [], [f11071, f297])).
fof(f11071, plain, (m2_filter_2(k19_filter_2(sK1, k2_xboole_0(sK2, sK3)), sK1) | v1_xboole_0(k2_xboole_0(sK2, sK3)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ m1_subset_1(sK2, k1_zfmisc_1(u1_struct_0(sK1))) | ~ m1_subset_1(sK3, k1_zfmisc_1(u1_struct_0(sK1)))), inference(superposition, [], [f925, f2078])).
fof(f925, plain, ! [X6, X8, X7] : (m2_filter_2(k19_filter_2(X6, k4_subset_1(u1_struct_0(X6), X7, X8)), X6) | v1_xboole_0(k4_subset_1(u1_struct_0(X6), X7, X8)) | ~ l3_lattices(X6) | ~ v10_lattices(X6) | v3_struct_0(X6) | ~ m1_subset_1(X8, k1_zfmisc_1(u1_struct_0(X6))) | ~ m1_subset_1(X7, k1_zfmisc_1(u1_struct_0(X6)))), inference(resolution, [], [f326, f331])).
fof(f2137, plain, (spl28_18 | spl28_80 | spl28_21 | ~ spl28_22 | ~ spl28_23), inference(avatar_split_clause, [], [f2132, f1319, f1315, f1311, f2134, f1285])).
fof(f1285, plain, (spl28_18 <=> ! [X0] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0))), introduced(avatar_definition, [new_symbols(naming, [spl28_18])])).
fof(f2132, plain, (! [X0] : ((k19_filter_2(sK1, sK3) = k3_filter_0(k1_lattice2(sK1), sK3)) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0)) | (spl28_21 | ~ spl28_22 | ~ spl28_23)), inference(forward_demodulation, [], [f2131, f1793])).
fof(f1793, plain, ((k3_filter_2(k1_lattice2(sK1), sK3) = k3_filter_0(k1_lattice2(sK1), sK3)) | (spl28_21 | ~ spl28_22 | ~ spl28_23)), inference(subsumption_resolution, [], [f1792, f1312])).
fof(f1792, plain, ((k3_filter_2(k1_lattice2(sK1), sK3) = k3_filter_0(k1_lattice2(sK1), sK3)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl28_22 | ~ spl28_23)), inference(subsumption_resolution, [], [f1791, f1316])).
fof(f1791, plain, ((k3_filter_2(k1_lattice2(sK1), sK3) = k3_filter_0(k1_lattice2(sK1), sK3)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ spl28_23), inference(subsumption_resolution, [], [f1790, f1320])).
fof(f1790, plain, ((k3_filter_2(k1_lattice2(sK1), sK3) = k3_filter_0(k1_lattice2(sK1), sK3)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(subsumption_resolution, [], [f1752, f300])).
fof(f1752, plain, ((k3_filter_2(k1_lattice2(sK1), sK3) = k3_filter_0(k1_lattice2(sK1), sK3)) | v1_xboole_0(sK3) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(resolution, [], [f1085, f437])).
fof(f2131, plain, ! [X0] : ((k19_filter_2(sK1, sK3) = k3_filter_2(k1_lattice2(sK1), sK3)) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0)), inference(forward_demodulation, [], [f2130, f1768])).
fof(f1768, plain, (sK3 = k8_filter_2(sK1, sK3)), inference(subsumption_resolution, [], [f1767, f295])).
fof(f1767, plain, ((sK3 = k8_filter_2(sK1, sK3)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1766, f296])).
fof(f1766, plain, ((sK3 = k8_filter_2(sK1, sK3)) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1746, f297])).
fof(f1746, plain, ((sK3 = k8_filter_2(sK1, sK3)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f1085, f322])).
fof(f2130, plain, ! [X0] : ((k3_filter_2(k1_lattice2(sK1), sK3) = k19_filter_2(sK1, k8_filter_2(sK1, sK3))) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f2129, f295])).
fof(f2129, plain, ! [X0] : ((k3_filter_2(k1_lattice2(sK1), sK3) = k19_filter_2(sK1, k8_filter_2(sK1, sK3))) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f2128, f296])).
fof(f2128, plain, ! [X0] : ((k3_filter_2(k1_lattice2(sK1), sK3) = k19_filter_2(sK1, k8_filter_2(sK1, sK3))) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f2127, f297])).
fof(f2127, plain, ! [X0] : ((k3_filter_2(k1_lattice2(sK1), sK3) = k19_filter_2(sK1, k8_filter_2(sK1, sK3))) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1741, f300])).
fof(f1741, plain, ! [X0] : ((k3_filter_2(k1_lattice2(sK1), sK3) = k19_filter_2(sK1, k8_filter_2(sK1, sK3))) | v1_xboole_0(sK3) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f1085, f454])).
fof(f454, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) | (k3_filter_2(k1_lattice2(X0), X2) = k19_filter_2(X0, k8_filter_2(X0, X2))) | v1_xboole_0(X2) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f227])).
fof(f2111, plain, ~ spl28_18, inference(avatar_contradiction_clause, [], [f2110])).
fof(f2110, plain, ($false | ~ spl28_18), inference(subsumption_resolution, [], [f2094, f300])).
fof(f2094, plain, (v1_xboole_0(sK3) | ~ spl28_18), inference(resolution, [], [f1286, f301])).
fof(f1286, plain, (! [X0] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0)) | ~ spl28_18), inference(avatar_component_clause, [], [f1285])).
fof(f1659, plain, spl28_61, inference(avatar_contradiction_clause, [], [f1658])).
fof(f1658, plain, ($false | spl28_61), inference(subsumption_resolution, [], [f1657, f480])).
fof(f1657, plain, (~ l1_lattices(sK1) | spl28_61), inference(resolution, [], [f1624, f346])).
fof(f346, plain, ! [X0] : (v1_funct_1(u1_lattices(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f166])).
fof(f1624, plain, (~ v1_funct_1(u1_lattices(sK1)) | spl28_61), inference(avatar_component_clause, [], [f1622])).
fof(f1622, plain, (spl28_61 <=> v1_funct_1(u1_lattices(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl28_61])])).
fof(f1656, plain, (~ spl28_61 | ~ spl28_62 | ~ spl28_63 | ~ spl28_64 | ~ spl28_65 | ~ spl28_66 | spl28_69), inference(avatar_split_clause, [], [f1620, f1654, f1642, f1638, f1634, f1630, f1626, f1622])).
fof(f1620, plain, ! [X17, X15, X16] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15) | ~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u2_lattices(sK1)) | ~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u1_lattices(sK1))), inference(superposition, [], [f397, f722])).
fof(f722, plain, (k1_lattice2(sK1) = g3_lattices(u1_struct_0(sK1), u1_lattices(sK1), u2_lattices(sK1))), inference(resolution, [], [f320, f297])).
fof(f320, plain, ! [X0] : (~ l3_lattices(X0) | (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ! [X0] : ((k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ! [X0] : (l3_lattices(X0) => (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', d2_lattice2)).
fof(f397, plain, ! [X4, X2, X0, X5, X3, X1] : (~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) | (X0 = X3) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f201])).
fof(f201, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(flattening, [], [f200])).
fof(f200, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | (~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1))), inference(ennf_transformation, [], [f68])).
fof(f68, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X2) & m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X1)) => ! [X3, X4, X5] : ((g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) => ((X2 = X5) & (X1 = X4) & (X0 = X3)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', free_g3_lattices)).
fof(f1438, plain, ~ spl28_21, inference(avatar_contradiction_clause, [], [f1437])).
fof(f1437, plain, ($false | ~ spl28_21), inference(subsumption_resolution, [], [f1433, f630])).
fof(f630, plain, sP0(sK1), inference(subsumption_resolution, [], [f629, f295])).
fof(f629, plain, (sP0(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f624, f296])).
fof(f624, plain, (sP0(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f396, f297])).
fof(f396, plain, ! [X0] : (~ l3_lattices(X0) | sP0(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f238])).
fof(f238, plain, ! [X0] : (sP0(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(definition_folding, [], [f199, e237])).
fof(f237, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(usedef, [], [e237])).
fof(e237, plain, ! [X0] : (sP0(X0) <=> (v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f199, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f198])).
fof(f198, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', fc6_lattice2)).
fof(f1433, plain, (~ sP0(sK1) | ~ spl28_21), inference(resolution, [], [f1313, f387])).
fof(f387, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f265])).
fof(f265, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(nnf_transformation, [], [f237])).
fof(f1313, plain, (v3_struct_0(k1_lattice2(sK1)) | ~ spl28_21), inference(avatar_component_clause, [], [f1311])).
fof(f1430, plain, spl28_22, inference(avatar_contradiction_clause, [], [f1429])).
fof(f1429, plain, ($false | spl28_22), inference(subsumption_resolution, [], [f1428, f630])).
fof(f1428, plain, (~ sP0(sK1) | spl28_22), inference(resolution, [], [f1317, f395])).
fof(f395, plain, ! [X0] : (v10_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f265])).
fof(f1317, plain, (~ v10_lattices(k1_lattice2(sK1)) | spl28_22), inference(avatar_component_clause, [], [f1315])).
fof(f1421, plain, (spl28_21 | ~ spl28_22 | ~ spl28_23 | spl28_32), inference(avatar_split_clause, [], [f1417, f1419, f1319, f1315, f1311])).
fof(f1417, plain, ! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, sK2)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, k3_filter_0(k1_lattice2(sK1), sK2)))) | ~ m1_subset_1(X22, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))) | v1_xboole_0(X22) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(subsumption_resolution, [], [f1400, f298])).
fof(f1400, plain, ! [X22] : ((k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, sK2)) = k3_filter_0(k1_lattice2(sK1), k4_subset_1(u1_struct_0(k1_lattice2(sK1)), X22, k3_filter_0(k1_lattice2(sK1), sK2)))) | v1_xboole_0(sK2) | ~ m1_subset_1(X22, k1_zfmisc_1(u1_struct_0(k1_lattice2(sK1)))) | v1_xboole_0(X22) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(resolution, [], [f457, f1081])).
fof(f457, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) | (k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, X2)) = k3_filter_0(X0, k4_subset_1(u1_struct_0(X0), X1, k3_filter_0(X0, X2)))) | v1_xboole_0(X2) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f230])).
fof(f1416, plain, spl28_23, inference(avatar_contradiction_clause, [], [f1415])).
fof(f1415, plain, ($false | spl28_23), inference(subsumption_resolution, [], [f1414, f297])).
fof(f1414, plain, (~ l3_lattices(sK1) | spl28_23), inference(resolution, [], [f1321, f328])).
fof(f1321, plain, (~ l3_lattices(k1_lattice2(sK1)) | spl28_23), inference(avatar_component_clause, [], [f1319])).
fof(f1391, plain, ~ spl28_13, inference(avatar_contradiction_clause, [], [f1390])).
fof(f1390, plain, ($false | ~ spl28_13), inference(subsumption_resolution, [], [f1389, f298])).
fof(f1389, plain, (v1_xboole_0(sK2) | ~ spl28_13), inference(resolution, [], [f670, f639])).
fof(f639, plain, ! [X3] : (r2_hidden(sK11(X3), X3) | v1_xboole_0(X3)), inference(resolution, [], [f450, f359])).
fof(f359, plain, ! [X0] : m1_subset_1(sK11(X0), X0), inference(cnf_transformation, [], [f258])).
fof(f258, plain, ! [X0] : m1_subset_1(sK11(X0), X0), inference(skolemisation, [status(esa), new_symbols(skolem, [sK11])], [f49, f257])).
fof(f257, plain, ! [X0] : (? [X1] : m1_subset_1(X1, X0) => m1_subset_1(sK11(X0), X0)), introduced(choice_axiom, [])).
fof(f49, plain, ! [X0] : ? [X1] : m1_subset_1(X1, X0), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', existence_m1_subset_1)).
fof(f450, plain, ! [X0, X1] : (~ m1_subset_1(X0, X1) | v1_xboole_0(X1) | r2_hidden(X0, X1)), inference(cnf_transformation, [], [f225])).
fof(f225, plain, ! [X0, X1] : (r2_hidden(X0, X1) | v1_xboole_0(X1) | ~ m1_subset_1(X0, X1)), inference(flattening, [], [f224])).
fof(f224, plain, ! [X0, X1] : ((r2_hidden(X0, X1) | v1_xboole_0(X1)) | ~ m1_subset_1(X0, X1)), inference(ennf_transformation, [], [f95])).
fof(f95, plain, ! [X0, X1] : (m1_subset_1(X0, X1) => (r2_hidden(X0, X1) | v1_xboole_0(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', t2_subset)).
fof(f670, plain, (! [X5] : ~ r2_hidden(X5, sK2) | ~ spl28_13), inference(avatar_component_clause, [], [f669])).
fof(f669, plain, (spl28_13 <=> ! [X5] : ~ r2_hidden(X5, sK2)), introduced(avatar_definition, [new_symbols(naming, [spl28_13])])).
fof(f1365, plain, (spl28_13 | ~ spl28_30), inference(avatar_split_clause, [], [f1278, f1357, f669])).
fof(f1278, plain, ! [X8] : (~ v1_xboole_0(u1_struct_0(k1_lattice2(sK1))) | ~ r2_hidden(X8, sK2)), inference(resolution, [], [f1081, f459])).
fof(f459, plain, ! [X2, X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(X2)) | ~ v1_xboole_0(X2) | ~ r2_hidden(X0, X1)), inference(cnf_transformation, [], [f233])).
fof(f233, plain, ! [X0, X1, X2] : (~ v1_xboole_0(X2) | ~ m1_subset_1(X1, k1_zfmisc_1(X2)) | ~ r2_hidden(X0, X1)), inference(ennf_transformation, [], [f100])).
fof(f100, plain, ! [X0, X1, X2] : ~ (v1_xboole_0(X2) & m1_subset_1(X1, k1_zfmisc_1(X2)) & r2_hidden(X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT315+1.p', t5_subset)).
fof(f1338, plain, (spl28_21 | ~ spl28_22 | ~ spl28_23 | spl28_26), inference(avatar_split_clause, [], [f1333, f1335, f1319, f1315, f1311])).
fof(f1333, plain, (m1_filter_0(k3_filter_0(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(subsumption_resolution, [], [f1269, f298])).
fof(f1269, plain, (m1_filter_0(k3_filter_0(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | v1_xboole_0(sK2) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(resolution, [], [f1081, f329])).
fof(f1326, plain, (spl28_21 | ~ spl28_22 | ~ spl28_23 | spl28_24), inference(avatar_split_clause, [], [f1309, f1323, f1319, f1315, f1311])).
fof(f1309, plain, ((k3_filter_2(k1_lattice2(sK1), sK2) = k3_filter_0(k1_lattice2(sK1), sK2)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(subsumption_resolution, [], [f1267, f298])).
fof(f1267, plain, ((k3_filter_2(k1_lattice2(sK1), sK2) = k3_filter_0(k1_lattice2(sK1), sK2)) | v1_xboole_0(sK2) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(resolution, [], [f1081, f437])).
fof(f1291, plain, (spl28_18 | spl28_19), inference(avatar_split_clause, [], [f1283, f1288, f1285])).
fof(f1283, plain, ! [X0] : ((k3_filter_2(k1_lattice2(sK1), sK2) = k19_filter_2(sK1, k8_filter_2(sK1, sK2))) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f1282, f295])).
fof(f1282, plain, ! [X0] : ((k3_filter_2(k1_lattice2(sK1), sK2) = k19_filter_2(sK1, k8_filter_2(sK1, sK2))) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1281, f296])).
fof(f1281, plain, ! [X0] : ((k3_filter_2(k1_lattice2(sK1), sK2) = k19_filter_2(sK1, k8_filter_2(sK1, sK2))) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1280, f297])).
fof(f1280, plain, ! [X0] : ((k3_filter_2(k1_lattice2(sK1), sK2) = k19_filter_2(sK1, k8_filter_2(sK1, sK2))) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1261, f298])).
fof(f1261, plain, ! [X0] : ((k3_filter_2(k1_lattice2(sK1), sK2) = k19_filter_2(sK1, k8_filter_2(sK1, sK2))) | v1_xboole_0(sK2) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(sK1))) | v1_xboole_0(X0) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f1081, f454])).
fof(f675, plain, (spl28_13 | ~ spl28_14), inference(avatar_split_clause, [], [f657, f672, f669])).
fof(f657, plain, ! [X5] : (~ v1_xboole_0(u1_struct_0(sK1)) | ~ r2_hidden(X5, sK2)), inference(resolution, [], [f459, f299])).
fof(f473, plain, (~ spl28_1 | ~ spl28_2), inference(avatar_split_clause, [], [f302, f470, f466])).
fof(f302, plain, (~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, k19_filter_2(sK1, sK3)))) | ~ r1_filter_2(u1_struct_0(sK1), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), sK2, sK3)), k19_filter_2(sK1, k4_subset_1(u1_struct_0(sK1), k19_filter_2(sK1, sK2), sK3)))), inference(cnf_transformation, [], [f242])).