fof(f7516, plain, $false, inference(avatar_sat_refutation, [], [f2006, f2033, f2052, f2095, f2109, f2175, f2201, f2209, f2252, f2275, f3030, f3041, f3091, f3157, f7055, f7063, f7131, f7176, f7295, f7475, f7508, f7512, f7515])).
fof(f7515, plain, spl38_202, inference(avatar_contradiction_clause, [], [f7514])).
fof(f7514, plain, ($false | spl38_202), inference(subsumption_resolution, [], [f7513, f412])).
fof(f412, plain, ~ v2_setfam_1(sK6), inference(cnf_transformation, [], [f340])).
fof(f340, plain, (~ r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), k11_waybel34(sK6), k10_waybel34(sK6)) & ~ v2_setfam_1(sK6)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6])], [f187, f339])).
fof(f339, plain, (? [X0] : (~ r3_yellow20(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0), k11_waybel34(X0), k10_waybel34(X0)) & ~ v2_setfam_1(X0)) => (~ r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), k11_waybel34(sK6), k10_waybel34(sK6)) & ~ v2_setfam_1(sK6))), introduced(choice_axiom, [])).
fof(f187, plain, ? [X0] : (~ r3_yellow20(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0), k11_waybel34(X0), k10_waybel34(X0)) & ~ v2_setfam_1(X0)), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : (~ v2_setfam_1(X0) => r3_yellow20(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0), k11_waybel34(X0), k10_waybel34(X0))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : (~ v2_setfam_1(X0) => r3_yellow20(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0), k11_waybel34(X0), k10_waybel34(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', t59_waybel34)).
fof(f7513, plain, (v2_setfam_1(sK6) | spl38_202), inference(resolution, [], [f7046, f3276])).
fof(f3276, plain, ! [X0] : (v3_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f3275, f466])).
fof(f466, plain, ! [X0] : (~ v1_xboole_0(X0) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f213])).
fof(f213, plain, ! [X0] : (~ v1_xboole_0(X0) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ! [X0] : (~ v2_setfam_1(X0) => ~ v1_xboole_0(X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', cc2_setfam_1)).
fof(f3275, plain, ! [X0] : (v3_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3274, f537])).
fof(f537, plain, ! [X0] : (~ v3_struct_0(k4_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f243])).
fof(f243, plain, ! [X0] : ((l2_altcat_1(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0))) | v1_xboole_0(X0)), inference(ennf_transformation, [], [f53])).
fof(f53, plain, ! [X0] : (~ v1_xboole_0(X0) => (l2_altcat_1(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', dt_k4_waybel34)).
fof(f3274, plain, ! [X0] : (v3_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3273, f538])).
fof(f538, plain, ! [X0] : (v2_altcat_1(k4_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f243])).
fof(f3273, plain, ! [X0] : (v3_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3272, f540])).
fof(f540, plain, ! [X0] : (v11_altcat_1(k4_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f243])).
fof(f3272, plain, ! [X0] : (v3_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3271, f541])).
fof(f541, plain, ! [X0] : (v12_altcat_1(k4_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f243])).
fof(f3271, plain, ! [X0] : (v3_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ v12_altcat_1(k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3270, f543])).
fof(f543, plain, ! [X0] : (l2_altcat_1(k4_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f243])).
fof(f3270, plain, ! [X0] : (v3_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ l2_altcat_1(k4_waybel34(X0)) | ~ v12_altcat_1(k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3269, f557])).
fof(f557, plain, ! [X0] : (~ v3_struct_0(k8_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f247])).
fof(f247, plain, ! [X0] : ((m1_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) & v3_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) & v6_altcat_1(k8_waybel34(X0)) & v2_altcat_1(k8_waybel34(X0)) & ~ v3_struct_0(k8_waybel34(X0))) | v1_xboole_0(X0)), inference(ennf_transformation, [], [f57])).
fof(f57, plain, ! [X0] : (~ v1_xboole_0(X0) => (m1_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) & v3_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) & v6_altcat_1(k8_waybel34(X0)) & v2_altcat_1(k8_waybel34(X0)) & ~ v3_struct_0(k8_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', dt_k8_waybel34)).
fof(f3269, plain, ! [X0] : (v3_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | v3_struct_0(k8_waybel34(X0)) | ~ l2_altcat_1(k4_waybel34(X0)) | ~ v12_altcat_1(k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3268, f558])).
fof(f558, plain, ! [X0] : (v2_altcat_1(k8_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f247])).
fof(f3268, plain, ! [X0] : (v3_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ v2_altcat_1(k8_waybel34(X0)) | v3_struct_0(k8_waybel34(X0)) | ~ l2_altcat_1(k4_waybel34(X0)) | ~ v12_altcat_1(k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3267, f561])).
fof(f561, plain, ! [X0] : (m1_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f247])).
fof(f3267, plain, ! [X0] : (~ m1_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) | v3_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ v2_altcat_1(k8_waybel34(X0)) | v3_struct_0(k8_waybel34(X0)) | ~ l2_altcat_1(k4_waybel34(X0)) | ~ v12_altcat_1(k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(resolution, [], [f1616, f560])).
fof(f560, plain, ! [X0] : (v3_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f247])).
fof(f1616, plain, ! [X4, X5] : (~ v3_altcat_2(k8_waybel34(X4), X5) | ~ m1_altcat_2(k8_waybel34(X4), X5) | v3_altcat_2(k10_waybel34(X4), X5) | ~ v2_altcat_1(k8_waybel34(X4)) | v3_struct_0(k8_waybel34(X4)) | ~ l2_altcat_1(X5) | ~ v12_altcat_1(X5) | ~ v11_altcat_1(X5) | ~ v2_altcat_1(X5) | v3_struct_0(X5) | v2_setfam_1(X4)), inference(subsumption_resolution, [], [f1615, f523])).
fof(f523, plain, ! [X0] : (~ v3_struct_0(k10_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f236])).
fof(f236, plain, ! [X0] : ((m1_altcat_2(k10_waybel34(X0), k8_waybel34(X0)) & v3_altcat_2(k10_waybel34(X0), k8_waybel34(X0)) & v6_altcat_1(k10_waybel34(X0)) & v2_altcat_1(k10_waybel34(X0)) & ~ v3_struct_0(k10_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f139])).
fof(f139, plain, ! [X0] : (~ v2_setfam_1(X0) => (m1_altcat_2(k10_waybel34(X0), k8_waybel34(X0)) & v3_altcat_2(k10_waybel34(X0), k8_waybel34(X0)) & v6_altcat_1(k10_waybel34(X0)) & v2_altcat_1(k10_waybel34(X0)) & ~ v3_struct_0(k10_waybel34(X0)))), inference(pure_predicate_removal, [], [f44])).
fof(f44, plain, ! [X0] : (~ v2_setfam_1(X0) => (m1_altcat_2(k10_waybel34(X0), k8_waybel34(X0)) & v3_altcat_2(k10_waybel34(X0), k8_waybel34(X0)) & v2_altcat_2(k10_waybel34(X0), k8_waybel34(X0)) & v6_altcat_1(k10_waybel34(X0)) & v2_altcat_1(k10_waybel34(X0)) & ~ v3_struct_0(k10_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', dt_k10_waybel34)).
fof(f1615, plain, ! [X4, X5] : (v3_altcat_2(k10_waybel34(X4), X5) | v3_struct_0(k10_waybel34(X4)) | ~ m1_altcat_2(k8_waybel34(X4), X5) | ~ v3_altcat_2(k8_waybel34(X4), X5) | ~ v2_altcat_1(k8_waybel34(X4)) | v3_struct_0(k8_waybel34(X4)) | ~ l2_altcat_1(X5) | ~ v12_altcat_1(X5) | ~ v11_altcat_1(X5) | ~ v2_altcat_1(X5) | v3_struct_0(X5) | v2_setfam_1(X4)), inference(subsumption_resolution, [], [f1614, f524])).
fof(f524, plain, ! [X0] : (v2_altcat_1(k10_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f236])).
fof(f1614, plain, ! [X4, X5] : (v3_altcat_2(k10_waybel34(X4), X5) | ~ v2_altcat_1(k10_waybel34(X4)) | v3_struct_0(k10_waybel34(X4)) | ~ m1_altcat_2(k8_waybel34(X4), X5) | ~ v3_altcat_2(k8_waybel34(X4), X5) | ~ v2_altcat_1(k8_waybel34(X4)) | v3_struct_0(k8_waybel34(X4)) | ~ l2_altcat_1(X5) | ~ v12_altcat_1(X5) | ~ v11_altcat_1(X5) | ~ v2_altcat_1(X5) | v3_struct_0(X5) | v2_setfam_1(X4)), inference(subsumption_resolution, [], [f1602, f527])).
fof(f527, plain, ! [X0] : (m1_altcat_2(k10_waybel34(X0), k8_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f236])).
fof(f1602, plain, ! [X4, X5] : (~ m1_altcat_2(k10_waybel34(X4), k8_waybel34(X4)) | v3_altcat_2(k10_waybel34(X4), X5) | ~ v2_altcat_1(k10_waybel34(X4)) | v3_struct_0(k10_waybel34(X4)) | ~ m1_altcat_2(k8_waybel34(X4), X5) | ~ v3_altcat_2(k8_waybel34(X4), X5) | ~ v2_altcat_1(k8_waybel34(X4)) | v3_struct_0(k8_waybel34(X4)) | ~ l2_altcat_1(X5) | ~ v12_altcat_1(X5) | ~ v11_altcat_1(X5) | ~ v2_altcat_1(X5) | v3_struct_0(X5) | v2_setfam_1(X4)), inference(resolution, [], [f732, f526])).
fof(f526, plain, ! [X0] : (v3_altcat_2(k10_waybel34(X0), k8_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f236])).
fof(f732, plain, ! [X2, X0, X1] : (~ v3_altcat_2(X2, X1) | ~ m1_altcat_2(X2, X1) | v3_altcat_2(X2, X0) | ~ v2_altcat_1(X2) | v3_struct_0(X2) | ~ m1_altcat_2(X1, X0) | ~ v3_altcat_2(X1, X0) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ l2_altcat_1(X0) | ~ v12_altcat_1(X0) | ~ v11_altcat_1(X0) | ~ v2_altcat_1(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f316])).
fof(f316, plain, ! [X0] : (! [X1] : (! [X2] : ((m1_altcat_2(X2, X0) & v3_altcat_2(X2, X0) & v2_altcat_1(X2) & ~ v3_struct_0(X2)) | ~ m1_altcat_2(X2, X1) | ~ v3_altcat_2(X2, X1) | ~ v2_altcat_1(X2) | v3_struct_0(X2)) | ~ m1_altcat_2(X1, X0) | ~ v3_altcat_2(X1, X0) | ~ v2_altcat_1(X1) | v3_struct_0(X1)) | ~ l2_altcat_1(X0) | ~ v12_altcat_1(X0) | ~ v11_altcat_1(X0) | ~ v2_altcat_1(X0) | v3_struct_0(X0)), inference(flattening, [], [f315])).
fof(f315, plain, ! [X0] : (! [X1] : (! [X2] : ((m1_altcat_2(X2, X0) & v3_altcat_2(X2, X0) & v2_altcat_1(X2) & ~ v3_struct_0(X2)) | (~ m1_altcat_2(X2, X1) | ~ v3_altcat_2(X2, X1) | ~ v2_altcat_1(X2) | v3_struct_0(X2))) | (~ m1_altcat_2(X1, X0) | ~ v3_altcat_2(X1, X0) | ~ v2_altcat_1(X1) | v3_struct_0(X1))) | (~ l2_altcat_1(X0) | ~ v12_altcat_1(X0) | ~ v11_altcat_1(X0) | ~ v2_altcat_1(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f125])).
fof(f125, plain, ! [X0] : ((l2_altcat_1(X0) & v12_altcat_1(X0) & v11_altcat_1(X0) & v2_altcat_1(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m1_altcat_2(X1, X0) & v3_altcat_2(X1, X0) & v2_altcat_1(X1) & ~ v3_struct_0(X1)) => ! [X2] : ((m1_altcat_2(X2, X1) & v3_altcat_2(X2, X1) & v2_altcat_1(X2) & ~ v3_struct_0(X2)) => (m1_altcat_2(X2, X0) & v3_altcat_2(X2, X0) & v2_altcat_1(X2) & ~ v3_struct_0(X2))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', t36_altcat_4)).
fof(f7046, plain, (~ v3_altcat_2(k10_waybel34(sK6), k4_waybel34(sK6)) | spl38_202), inference(avatar_component_clause, [], [f7044])).
fof(f7044, plain, (spl38_202 <=> v3_altcat_2(k10_waybel34(sK6), k4_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_202])])).
fof(f7512, plain, ~ spl38_204, inference(avatar_contradiction_clause, [], [f7511])).
fof(f7511, plain, ($false | ~ spl38_204), inference(subsumption_resolution, [], [f7510, f412])).
fof(f7510, plain, (v2_setfam_1(sK6) | ~ spl38_204), inference(resolution, [], [f7054, f523])).
fof(f7054, plain, (v3_struct_0(k10_waybel34(sK6)) | ~ spl38_204), inference(avatar_component_clause, [], [f7052])).
fof(f7052, plain, (spl38_204 <=> v3_struct_0(k10_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_204])])).
fof(f7508, plain, spl38_203, inference(avatar_contradiction_clause, [], [f7507])).
fof(f7507, plain, ($false | spl38_203), inference(subsumption_resolution, [], [f7506, f412])).
fof(f7506, plain, (v2_setfam_1(sK6) | spl38_203), inference(resolution, [], [f7050, f524])).
fof(f7050, plain, (~ v2_altcat_1(k10_waybel34(sK6)) | spl38_203), inference(avatar_component_clause, [], [f7048])).
fof(f7048, plain, (spl38_203 <=> v2_altcat_1(k10_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_203])])).
fof(f7475, plain, spl38_201, inference(avatar_contradiction_clause, [], [f7474])).
fof(f7474, plain, ($false | spl38_201), inference(subsumption_resolution, [], [f7473, f412])).
fof(f7473, plain, (v2_setfam_1(sK6) | spl38_201), inference(resolution, [], [f7042, f3290])).
fof(f3290, plain, ! [X0] : (m1_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f3289, f466])).
fof(f3289, plain, ! [X0] : (m1_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3288, f537])).
fof(f3288, plain, ! [X0] : (m1_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3287, f538])).
fof(f3287, plain, ! [X0] : (m1_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3286, f540])).
fof(f3286, plain, ! [X0] : (m1_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3285, f541])).
fof(f3285, plain, ! [X0] : (m1_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ v12_altcat_1(k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3284, f543])).
fof(f3284, plain, ! [X0] : (m1_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ l2_altcat_1(k4_waybel34(X0)) | ~ v12_altcat_1(k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3283, f557])).
fof(f3283, plain, ! [X0] : (m1_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | v3_struct_0(k8_waybel34(X0)) | ~ l2_altcat_1(k4_waybel34(X0)) | ~ v12_altcat_1(k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3282, f558])).
fof(f3282, plain, ! [X0] : (m1_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ v2_altcat_1(k8_waybel34(X0)) | v3_struct_0(k8_waybel34(X0)) | ~ l2_altcat_1(k4_waybel34(X0)) | ~ v12_altcat_1(k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(subsumption_resolution, [], [f3281, f561])).
fof(f3281, plain, ! [X0] : (~ m1_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) | m1_altcat_2(k10_waybel34(X0), k4_waybel34(X0)) | ~ v2_altcat_1(k8_waybel34(X0)) | v3_struct_0(k8_waybel34(X0)) | ~ l2_altcat_1(k4_waybel34(X0)) | ~ v12_altcat_1(k4_waybel34(X0)) | ~ v11_altcat_1(k4_waybel34(X0)) | ~ v2_altcat_1(k4_waybel34(X0)) | v3_struct_0(k4_waybel34(X0)) | v2_setfam_1(X0) | v1_xboole_0(X0)), inference(resolution, [], [f1644, f560])).
fof(f1644, plain, ! [X4, X5] : (~ v3_altcat_2(k8_waybel34(X4), X5) | ~ m1_altcat_2(k8_waybel34(X4), X5) | m1_altcat_2(k10_waybel34(X4), X5) | ~ v2_altcat_1(k8_waybel34(X4)) | v3_struct_0(k8_waybel34(X4)) | ~ l2_altcat_1(X5) | ~ v12_altcat_1(X5) | ~ v11_altcat_1(X5) | ~ v2_altcat_1(X5) | v3_struct_0(X5) | v2_setfam_1(X4)), inference(subsumption_resolution, [], [f1643, f523])).
fof(f1643, plain, ! [X4, X5] : (m1_altcat_2(k10_waybel34(X4), X5) | v3_struct_0(k10_waybel34(X4)) | ~ m1_altcat_2(k8_waybel34(X4), X5) | ~ v3_altcat_2(k8_waybel34(X4), X5) | ~ v2_altcat_1(k8_waybel34(X4)) | v3_struct_0(k8_waybel34(X4)) | ~ l2_altcat_1(X5) | ~ v12_altcat_1(X5) | ~ v11_altcat_1(X5) | ~ v2_altcat_1(X5) | v3_struct_0(X5) | v2_setfam_1(X4)), inference(subsumption_resolution, [], [f1642, f524])).
fof(f1642, plain, ! [X4, X5] : (m1_altcat_2(k10_waybel34(X4), X5) | ~ v2_altcat_1(k10_waybel34(X4)) | v3_struct_0(k10_waybel34(X4)) | ~ m1_altcat_2(k8_waybel34(X4), X5) | ~ v3_altcat_2(k8_waybel34(X4), X5) | ~ v2_altcat_1(k8_waybel34(X4)) | v3_struct_0(k8_waybel34(X4)) | ~ l2_altcat_1(X5) | ~ v12_altcat_1(X5) | ~ v11_altcat_1(X5) | ~ v2_altcat_1(X5) | v3_struct_0(X5) | v2_setfam_1(X4)), inference(subsumption_resolution, [], [f1630, f527])).
fof(f1630, plain, ! [X4, X5] : (~ m1_altcat_2(k10_waybel34(X4), k8_waybel34(X4)) | m1_altcat_2(k10_waybel34(X4), X5) | ~ v2_altcat_1(k10_waybel34(X4)) | v3_struct_0(k10_waybel34(X4)) | ~ m1_altcat_2(k8_waybel34(X4), X5) | ~ v3_altcat_2(k8_waybel34(X4), X5) | ~ v2_altcat_1(k8_waybel34(X4)) | v3_struct_0(k8_waybel34(X4)) | ~ l2_altcat_1(X5) | ~ v12_altcat_1(X5) | ~ v11_altcat_1(X5) | ~ v2_altcat_1(X5) | v3_struct_0(X5) | v2_setfam_1(X4)), inference(resolution, [], [f733, f526])).
fof(f733, plain, ! [X2, X0, X1] : (~ v3_altcat_2(X2, X1) | ~ m1_altcat_2(X2, X1) | m1_altcat_2(X2, X0) | ~ v2_altcat_1(X2) | v3_struct_0(X2) | ~ m1_altcat_2(X1, X0) | ~ v3_altcat_2(X1, X0) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ l2_altcat_1(X0) | ~ v12_altcat_1(X0) | ~ v11_altcat_1(X0) | ~ v2_altcat_1(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f316])).
fof(f7042, plain, (~ m1_altcat_2(k10_waybel34(sK6), k4_waybel34(sK6)) | spl38_201), inference(avatar_component_clause, [], [f7040])).
fof(f7040, plain, (spl38_201 <=> m1_altcat_2(k10_waybel34(sK6), k4_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_201])])).
fof(f7295, plain, (spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18 | spl38_198), inference(avatar_contradiction_clause, [], [f7294])).
fof(f7294, plain, ($false | (spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18 | spl38_198)), inference(subsumption_resolution, [], [f7293, f412])).
fof(f7293, plain, (v2_setfam_1(sK6) | (spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18 | spl38_198)), inference(subsumption_resolution, [], [f7292, f1943])).
fof(f1943, plain, (~ v3_struct_0(k5_waybel34(sK6)) | spl38_14), inference(avatar_component_clause, [], [f1942])).
fof(f1942, plain, (spl38_14 <=> v3_struct_0(k5_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_14])])).
fof(f7292, plain, (v3_struct_0(k5_waybel34(sK6)) | v2_setfam_1(sK6) | (~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18 | spl38_198)), inference(subsumption_resolution, [], [f7291, f1947])).
fof(f1947, plain, (v2_altcat_1(k5_waybel34(sK6)) | ~ spl38_15), inference(avatar_component_clause, [], [f1946])).
fof(f1946, plain, (spl38_15 <=> v2_altcat_1(k5_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_15])])).
fof(f7291, plain, (~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | v2_setfam_1(sK6) | (~ spl38_16 | ~ spl38_17 | ~ spl38_18 | spl38_198)), inference(subsumption_resolution, [], [f7290, f1951])).
fof(f1951, plain, (v11_altcat_1(k5_waybel34(sK6)) | ~ spl38_16), inference(avatar_component_clause, [], [f1950])).
fof(f1950, plain, (spl38_16 <=> v11_altcat_1(k5_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_16])])).
fof(f7290, plain, (~ v11_altcat_1(k5_waybel34(sK6)) | ~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | v2_setfam_1(sK6) | (~ spl38_17 | ~ spl38_18 | spl38_198)), inference(subsumption_resolution, [], [f7289, f1955])).
fof(f1955, plain, (v12_altcat_1(k5_waybel34(sK6)) | ~ spl38_17), inference(avatar_component_clause, [], [f1954])).
fof(f1954, plain, (spl38_17 <=> v12_altcat_1(k5_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_17])])).
fof(f7289, plain, (~ v12_altcat_1(k5_waybel34(sK6)) | ~ v11_altcat_1(k5_waybel34(sK6)) | ~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | v2_setfam_1(sK6) | (~ spl38_18 | spl38_198)), inference(subsumption_resolution, [], [f7288, f1959])).
fof(f1959, plain, (l2_altcat_1(k5_waybel34(sK6)) | ~ spl38_18), inference(avatar_component_clause, [], [f1958])).
fof(f1958, plain, (spl38_18 <=> l2_altcat_1(k5_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_18])])).
fof(f7288, plain, (~ l2_altcat_1(k5_waybel34(sK6)) | ~ v12_altcat_1(k5_waybel34(sK6)) | ~ v11_altcat_1(k5_waybel34(sK6)) | ~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | v2_setfam_1(sK6) | spl38_198), inference(resolution, [], [f7030, f3088])).
fof(f3088, plain, ! [X0] : (v3_altcat_2(k11_waybel34(X0), k5_waybel34(X0)) | ~ l2_altcat_1(k5_waybel34(X0)) | ~ v12_altcat_1(k5_waybel34(X0)) | ~ v11_altcat_1(k5_waybel34(X0)) | ~ v2_altcat_1(k5_waybel34(X0)) | v3_struct_0(k5_waybel34(X0)) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f3087, f566])).
fof(f566, plain, ! [X0] : (m1_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f248])).
fof(f248, plain, ! [X0] : ((m1_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) & v3_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) & v6_altcat_1(k9_waybel34(X0)) & v2_altcat_1(k9_waybel34(X0)) & ~ v3_struct_0(k9_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f58])).
fof(f58, plain, ! [X0] : (~ v2_setfam_1(X0) => (m1_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) & v3_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) & v6_altcat_1(k9_waybel34(X0)) & v2_altcat_1(k9_waybel34(X0)) & ~ v3_struct_0(k9_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', dt_k9_waybel34)).
fof(f3087, plain, ! [X0] : (~ m1_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) | v3_altcat_2(k11_waybel34(X0), k5_waybel34(X0)) | ~ l2_altcat_1(k5_waybel34(X0)) | ~ v12_altcat_1(k5_waybel34(X0)) | ~ v11_altcat_1(k5_waybel34(X0)) | ~ v2_altcat_1(k5_waybel34(X0)) | v3_struct_0(k5_waybel34(X0)) | v2_setfam_1(X0)), inference(duplicate_literal_removal, [], [f3086])).
fof(f3086, plain, ! [X0] : (~ m1_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) | v3_altcat_2(k11_waybel34(X0), k5_waybel34(X0)) | ~ l2_altcat_1(k5_waybel34(X0)) | ~ v12_altcat_1(k5_waybel34(X0)) | ~ v11_altcat_1(k5_waybel34(X0)) | ~ v2_altcat_1(k5_waybel34(X0)) | v3_struct_0(k5_waybel34(X0)) | v2_setfam_1(X0) | v2_setfam_1(X0)), inference(resolution, [], [f1611, f565])).
fof(f565, plain, ! [X0] : (v3_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f248])).
fof(f1611, plain, ! [X0, X1] : (~ v3_altcat_2(k9_waybel34(X0), X1) | ~ m1_altcat_2(k9_waybel34(X0), X1) | v3_altcat_2(k11_waybel34(X0), X1) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f1610, f562])).
fof(f562, plain, ! [X0] : (~ v3_struct_0(k9_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f248])).
fof(f1610, plain, ! [X0, X1] : (v3_altcat_2(k11_waybel34(X0), X1) | ~ m1_altcat_2(k9_waybel34(X0), X1) | ~ v3_altcat_2(k9_waybel34(X0), X1) | v3_struct_0(k9_waybel34(X0)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f1609, f563])).
fof(f563, plain, ! [X0] : (v2_altcat_1(k9_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f248])).
fof(f1609, plain, ! [X0, X1] : (v3_altcat_2(k11_waybel34(X0), X1) | ~ m1_altcat_2(k9_waybel34(X0), X1) | ~ v3_altcat_2(k9_waybel34(X0), X1) | ~ v2_altcat_1(k9_waybel34(X0)) | v3_struct_0(k9_waybel34(X0)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f1608, f528])).
fof(f528, plain, ! [X0] : (~ v3_struct_0(k11_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f237])).
fof(f237, plain, ! [X0] : ((m1_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) & v3_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) & v6_altcat_1(k11_waybel34(X0)) & v2_altcat_1(k11_waybel34(X0)) & ~ v3_struct_0(k11_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f137])).
fof(f137, plain, ! [X0] : (~ v2_setfam_1(X0) => (m1_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) & v3_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) & v6_altcat_1(k11_waybel34(X0)) & v2_altcat_1(k11_waybel34(X0)) & ~ v3_struct_0(k11_waybel34(X0)))), inference(pure_predicate_removal, [], [f45])).
fof(f45, plain, ! [X0] : (~ v2_setfam_1(X0) => (m1_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) & v3_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) & v2_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) & v6_altcat_1(k11_waybel34(X0)) & v2_altcat_1(k11_waybel34(X0)) & ~ v3_struct_0(k11_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', dt_k11_waybel34)).
fof(f1608, plain, ! [X0, X1] : (v3_altcat_2(k11_waybel34(X0), X1) | v3_struct_0(k11_waybel34(X0)) | ~ m1_altcat_2(k9_waybel34(X0), X1) | ~ v3_altcat_2(k9_waybel34(X0), X1) | ~ v2_altcat_1(k9_waybel34(X0)) | v3_struct_0(k9_waybel34(X0)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f1607, f529])).
fof(f529, plain, ! [X0] : (v2_altcat_1(k11_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f237])).
fof(f1607, plain, ! [X0, X1] : (v3_altcat_2(k11_waybel34(X0), X1) | ~ v2_altcat_1(k11_waybel34(X0)) | v3_struct_0(k11_waybel34(X0)) | ~ m1_altcat_2(k9_waybel34(X0), X1) | ~ v3_altcat_2(k9_waybel34(X0), X1) | ~ v2_altcat_1(k9_waybel34(X0)) | v3_struct_0(k9_waybel34(X0)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f1600, f532])).
fof(f532, plain, ! [X0] : (m1_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f237])).
fof(f1600, plain, ! [X0, X1] : (~ m1_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) | v3_altcat_2(k11_waybel34(X0), X1) | ~ v2_altcat_1(k11_waybel34(X0)) | v3_struct_0(k11_waybel34(X0)) | ~ m1_altcat_2(k9_waybel34(X0), X1) | ~ v3_altcat_2(k9_waybel34(X0), X1) | ~ v2_altcat_1(k9_waybel34(X0)) | v3_struct_0(k9_waybel34(X0)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(resolution, [], [f732, f531])).
fof(f531, plain, ! [X0] : (v3_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f237])).
fof(f7030, plain, (~ v3_altcat_2(k11_waybel34(sK6), k5_waybel34(sK6)) | spl38_198), inference(avatar_component_clause, [], [f7028])).
fof(f7028, plain, (spl38_198 <=> v3_altcat_2(k11_waybel34(sK6), k5_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_198])])).
fof(f7176, plain, ~ spl38_200, inference(avatar_contradiction_clause, [], [f7175])).
fof(f7175, plain, ($false | ~ spl38_200), inference(subsumption_resolution, [], [f7174, f412])).
fof(f7174, plain, (v2_setfam_1(sK6) | ~ spl38_200), inference(resolution, [], [f7038, f528])).
fof(f7038, plain, (v3_struct_0(k11_waybel34(sK6)) | ~ spl38_200), inference(avatar_component_clause, [], [f7036])).
fof(f7036, plain, (spl38_200 <=> v3_struct_0(k11_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_200])])).
fof(f7131, plain, spl38_199, inference(avatar_contradiction_clause, [], [f7130])).
fof(f7130, plain, ($false | spl38_199), inference(subsumption_resolution, [], [f7129, f412])).
fof(f7129, plain, (v2_setfam_1(sK6) | spl38_199), inference(resolution, [], [f7034, f529])).
fof(f7034, plain, (~ v2_altcat_1(k11_waybel34(sK6)) | spl38_199), inference(avatar_component_clause, [], [f7032])).
fof(f7032, plain, (spl38_199 <=> v2_altcat_1(k11_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_199])])).
fof(f7063, plain, (spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18 | spl38_197), inference(avatar_contradiction_clause, [], [f7062])).
fof(f7062, plain, ($false | (spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18 | spl38_197)), inference(subsumption_resolution, [], [f7061, f412])).
fof(f7061, plain, (v2_setfam_1(sK6) | (spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18 | spl38_197)), inference(subsumption_resolution, [], [f7060, f1943])).
fof(f7060, plain, (v3_struct_0(k5_waybel34(sK6)) | v2_setfam_1(sK6) | (~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18 | spl38_197)), inference(subsumption_resolution, [], [f7059, f1947])).
fof(f7059, plain, (~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | v2_setfam_1(sK6) | (~ spl38_16 | ~ spl38_17 | ~ spl38_18 | spl38_197)), inference(subsumption_resolution, [], [f7058, f1951])).
fof(f7058, plain, (~ v11_altcat_1(k5_waybel34(sK6)) | ~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | v2_setfam_1(sK6) | (~ spl38_17 | ~ spl38_18 | spl38_197)), inference(subsumption_resolution, [], [f7057, f1955])).
fof(f7057, plain, (~ v12_altcat_1(k5_waybel34(sK6)) | ~ v11_altcat_1(k5_waybel34(sK6)) | ~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | v2_setfam_1(sK6) | (~ spl38_18 | spl38_197)), inference(subsumption_resolution, [], [f7056, f1959])).
fof(f7056, plain, (~ l2_altcat_1(k5_waybel34(sK6)) | ~ v12_altcat_1(k5_waybel34(sK6)) | ~ v11_altcat_1(k5_waybel34(sK6)) | ~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | v2_setfam_1(sK6) | spl38_197), inference(resolution, [], [f7026, f3154])).
fof(f3154, plain, ! [X0] : (m1_altcat_2(k11_waybel34(X0), k5_waybel34(X0)) | ~ l2_altcat_1(k5_waybel34(X0)) | ~ v12_altcat_1(k5_waybel34(X0)) | ~ v11_altcat_1(k5_waybel34(X0)) | ~ v2_altcat_1(k5_waybel34(X0)) | v3_struct_0(k5_waybel34(X0)) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f3153, f566])).
fof(f3153, plain, ! [X0] : (~ m1_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) | m1_altcat_2(k11_waybel34(X0), k5_waybel34(X0)) | ~ l2_altcat_1(k5_waybel34(X0)) | ~ v12_altcat_1(k5_waybel34(X0)) | ~ v11_altcat_1(k5_waybel34(X0)) | ~ v2_altcat_1(k5_waybel34(X0)) | v3_struct_0(k5_waybel34(X0)) | v2_setfam_1(X0)), inference(duplicate_literal_removal, [], [f3152])).
fof(f3152, plain, ! [X0] : (~ m1_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) | m1_altcat_2(k11_waybel34(X0), k5_waybel34(X0)) | ~ l2_altcat_1(k5_waybel34(X0)) | ~ v12_altcat_1(k5_waybel34(X0)) | ~ v11_altcat_1(k5_waybel34(X0)) | ~ v2_altcat_1(k5_waybel34(X0)) | v3_struct_0(k5_waybel34(X0)) | v2_setfam_1(X0) | v2_setfam_1(X0)), inference(resolution, [], [f1639, f565])).
fof(f1639, plain, ! [X0, X1] : (~ v3_altcat_2(k9_waybel34(X0), X1) | ~ m1_altcat_2(k9_waybel34(X0), X1) | m1_altcat_2(k11_waybel34(X0), X1) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f1638, f562])).
fof(f1638, plain, ! [X0, X1] : (m1_altcat_2(k11_waybel34(X0), X1) | ~ m1_altcat_2(k9_waybel34(X0), X1) | ~ v3_altcat_2(k9_waybel34(X0), X1) | v3_struct_0(k9_waybel34(X0)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f1637, f563])).
fof(f1637, plain, ! [X0, X1] : (m1_altcat_2(k11_waybel34(X0), X1) | ~ m1_altcat_2(k9_waybel34(X0), X1) | ~ v3_altcat_2(k9_waybel34(X0), X1) | ~ v2_altcat_1(k9_waybel34(X0)) | v3_struct_0(k9_waybel34(X0)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f1636, f528])).
fof(f1636, plain, ! [X0, X1] : (m1_altcat_2(k11_waybel34(X0), X1) | v3_struct_0(k11_waybel34(X0)) | ~ m1_altcat_2(k9_waybel34(X0), X1) | ~ v3_altcat_2(k9_waybel34(X0), X1) | ~ v2_altcat_1(k9_waybel34(X0)) | v3_struct_0(k9_waybel34(X0)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f1635, f529])).
fof(f1635, plain, ! [X0, X1] : (m1_altcat_2(k11_waybel34(X0), X1) | ~ v2_altcat_1(k11_waybel34(X0)) | v3_struct_0(k11_waybel34(X0)) | ~ m1_altcat_2(k9_waybel34(X0), X1) | ~ v3_altcat_2(k9_waybel34(X0), X1) | ~ v2_altcat_1(k9_waybel34(X0)) | v3_struct_0(k9_waybel34(X0)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(subsumption_resolution, [], [f1628, f532])).
fof(f1628, plain, ! [X0, X1] : (~ m1_altcat_2(k11_waybel34(X0), k9_waybel34(X0)) | m1_altcat_2(k11_waybel34(X0), X1) | ~ v2_altcat_1(k11_waybel34(X0)) | v3_struct_0(k11_waybel34(X0)) | ~ m1_altcat_2(k9_waybel34(X0), X1) | ~ v3_altcat_2(k9_waybel34(X0), X1) | ~ v2_altcat_1(k9_waybel34(X0)) | v3_struct_0(k9_waybel34(X0)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | v2_setfam_1(X0)), inference(resolution, [], [f733, f531])).
fof(f7026, plain, (~ m1_altcat_2(k11_waybel34(sK6), k5_waybel34(sK6)) | spl38_197), inference(avatar_component_clause, [], [f7024])).
fof(f7024, plain, (spl38_197 <=> m1_altcat_2(k11_waybel34(sK6), k5_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_197])])).
fof(f7055, plain, (~ spl38_197 | ~ spl38_198 | ~ spl38_199 | spl38_200 | ~ spl38_201 | ~ spl38_202 | ~ spl38_203 | spl38_204 | ~ spl38_22), inference(avatar_split_clause, [], [f7022, f1974, f7052, f7048, f7044, f7040, f7036, f7032, f7028, f7024])).
fof(f1974, plain, (spl38_22 <=> ! [X1, X0] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | v3_struct_0(X1) | ~ v2_altcat_1(X1) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | v3_struct_0(X0) | ~ v2_altcat_1(X0) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl38_22])])).
fof(f7022, plain, (v3_struct_0(k10_waybel34(sK6)) | ~ v2_altcat_1(k10_waybel34(sK6)) | ~ v3_altcat_2(k10_waybel34(sK6), k4_waybel34(sK6)) | ~ m1_altcat_2(k10_waybel34(sK6), k4_waybel34(sK6)) | v3_struct_0(k11_waybel34(sK6)) | ~ v2_altcat_1(k11_waybel34(sK6)) | ~ v3_altcat_2(k11_waybel34(sK6), k5_waybel34(sK6)) | ~ m1_altcat_2(k11_waybel34(sK6), k5_waybel34(sK6)) | ~ spl38_22), inference(subsumption_resolution, [], [f7021, f412])).
fof(f7021, plain, (v3_struct_0(k10_waybel34(sK6)) | ~ v2_altcat_1(k10_waybel34(sK6)) | ~ v3_altcat_2(k10_waybel34(sK6), k4_waybel34(sK6)) | ~ m1_altcat_2(k10_waybel34(sK6), k4_waybel34(sK6)) | v3_struct_0(k11_waybel34(sK6)) | ~ v2_altcat_1(k11_waybel34(sK6)) | ~ v3_altcat_2(k11_waybel34(sK6), k5_waybel34(sK6)) | ~ m1_altcat_2(k11_waybel34(sK6), k5_waybel34(sK6)) | v2_setfam_1(sK6) | ~ spl38_22), inference(subsumption_resolution, [], [f7020, f413])).
fof(f413, plain, ~ r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), k11_waybel34(sK6), k10_waybel34(sK6)), inference(cnf_transformation, [], [f340])).
fof(f7020, plain, (v3_struct_0(k10_waybel34(sK6)) | ~ v2_altcat_1(k10_waybel34(sK6)) | ~ v3_altcat_2(k10_waybel34(sK6), k4_waybel34(sK6)) | ~ m1_altcat_2(k10_waybel34(sK6), k4_waybel34(sK6)) | v3_struct_0(k11_waybel34(sK6)) | ~ v2_altcat_1(k11_waybel34(sK6)) | ~ v3_altcat_2(k11_waybel34(sK6), k5_waybel34(sK6)) | ~ m1_altcat_2(k11_waybel34(sK6), k5_waybel34(sK6)) | r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), k11_waybel34(sK6), k10_waybel34(sK6)) | v2_setfam_1(sK6) | ~ spl38_22), inference(resolution, [], [f1975, f737])).
fof(f737, plain, ! [X0] : (r3_yellow20(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0), k10_waybel34(X0), k11_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f322])).
fof(f322, plain, ! [X0] : (r3_yellow20(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0), k10_waybel34(X0), k11_waybel34(X0)) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f129])).
fof(f129, plain, ! [X0] : (~ v2_setfam_1(X0) => r3_yellow20(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0), k10_waybel34(X0), k11_waybel34(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', t58_waybel34)).
fof(f1975, plain, (! [X0, X1] : (~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | v3_struct_0(X1) | ~ v2_altcat_1(X1) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | v3_struct_0(X0) | ~ v2_altcat_1(X0) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1)) | ~ spl38_22), inference(avatar_component_clause, [], [f1974])).
fof(f3157, plain, spl38_21, inference(avatar_contradiction_clause, [], [f3156])).
fof(f3156, plain, ($false | spl38_21), inference(subsumption_resolution, [], [f3155, f412])).
fof(f3155, plain, (v2_setfam_1(sK6) | spl38_21), inference(resolution, [], [f1972, f651])).
fof(f651, plain, ! [X0] : (v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f290])).
fof(f290, plain, ! [X0] : ((v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v14_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v11_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f162])).
fof(f162, plain, ! [X0] : (~ v2_setfam_1(X0) => (v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v14_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v11_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)))), inference(pure_predicate_removal, [], [f150])).
fof(f150, plain, ! [X0] : (~ v2_setfam_1(X0) => (v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v14_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v11_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v8_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)))), inference(pure_predicate_removal, [], [f146])).
fof(f146, plain, ! [X0] : (~ v2_setfam_1(X0) => (v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v14_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v12_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v11_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v8_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)))), inference(pure_predicate_removal, [], [f98])).
fof(f98, plain, ! [X0] : (~ v2_setfam_1(X0) => (v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v14_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v12_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v11_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v8_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v6_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', fc7_waybel34)).
fof(f1972, plain, (~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | spl38_21), inference(avatar_component_clause, [], [f1970])).
fof(f1970, plain, (spl38_21 <=> v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_21])])).
fof(f3091, plain, spl38_20, inference(avatar_contradiction_clause, [], [f3090])).
fof(f3090, plain, ($false | spl38_20), inference(subsumption_resolution, [], [f3089, f412])).
fof(f3089, plain, (v2_setfam_1(sK6) | spl38_20), inference(resolution, [], [f1968, f553])).
fof(f553, plain, ! [X0] : (m2_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f245])).
fof(f245, plain, ! [X0] : ((m2_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f55])).
fof(f55, plain, ! [X0] : (~ v2_setfam_1(X0) => (m2_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', dt_k6_waybel34)).
fof(f1968, plain, (~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | spl38_20), inference(avatar_component_clause, [], [f1966])).
fof(f1966, plain, (spl38_20 <=> m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_20])])).
fof(f3041, plain, (~ spl38_19 | ~ spl38_20 | ~ spl38_21 | spl38_22 | spl38_9 | ~ spl38_10 | ~ spl38_11 | ~ spl38_12 | ~ spl38_13 | spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18), inference(avatar_split_clause, [], [f3040, f1958, f1954, f1950, f1946, f1942, f1938, f1934, f1930, f1926, f1922, f1974, f1970, f1966, f1962])).
fof(f1962, plain, (spl38_19 <=> v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_19])])).
fof(f1922, plain, (spl38_9 <=> v3_struct_0(k4_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_9])])).
fof(f1926, plain, (spl38_10 <=> v2_altcat_1(k4_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_10])])).
fof(f1930, plain, (spl38_11 <=> v11_altcat_1(k4_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_11])])).
fof(f1934, plain, (spl38_12 <=> v12_altcat_1(k4_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_12])])).
fof(f1938, plain, (spl38_13 <=> l2_altcat_1(k4_waybel34(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl38_13])])).
fof(f3040, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6))) | (spl38_9 | ~ spl38_10 | ~ spl38_11 | ~ spl38_12 | ~ spl38_13 | spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18)), inference(subsumption_resolution, [], [f3039, f1923])).
fof(f1923, plain, (~ v3_struct_0(k4_waybel34(sK6)) | spl38_9), inference(avatar_component_clause, [], [f1922])).
fof(f3039, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | v3_struct_0(k4_waybel34(sK6))) | (~ spl38_10 | ~ spl38_11 | ~ spl38_12 | ~ spl38_13 | spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18)), inference(subsumption_resolution, [], [f3038, f1927])).
fof(f1927, plain, (v2_altcat_1(k4_waybel34(sK6)) | ~ spl38_10), inference(avatar_component_clause, [], [f1926])).
fof(f3038, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v2_altcat_1(k4_waybel34(sK6)) | v3_struct_0(k4_waybel34(sK6))) | (~ spl38_11 | ~ spl38_12 | ~ spl38_13 | spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18)), inference(subsumption_resolution, [], [f3037, f1931])).
fof(f1931, plain, (v11_altcat_1(k4_waybel34(sK6)) | ~ spl38_11), inference(avatar_component_clause, [], [f1930])).
fof(f3037, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v11_altcat_1(k4_waybel34(sK6)) | ~ v2_altcat_1(k4_waybel34(sK6)) | v3_struct_0(k4_waybel34(sK6))) | (~ spl38_12 | ~ spl38_13 | spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18)), inference(subsumption_resolution, [], [f3036, f1935])).
fof(f1935, plain, (v12_altcat_1(k4_waybel34(sK6)) | ~ spl38_12), inference(avatar_component_clause, [], [f1934])).
fof(f3036, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v12_altcat_1(k4_waybel34(sK6)) | ~ v11_altcat_1(k4_waybel34(sK6)) | ~ v2_altcat_1(k4_waybel34(sK6)) | v3_struct_0(k4_waybel34(sK6))) | (~ spl38_13 | spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18)), inference(subsumption_resolution, [], [f3035, f1939])).
fof(f1939, plain, (l2_altcat_1(k4_waybel34(sK6)) | ~ spl38_13), inference(avatar_component_clause, [], [f1938])).
fof(f3035, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ l2_altcat_1(k4_waybel34(sK6)) | ~ v12_altcat_1(k4_waybel34(sK6)) | ~ v11_altcat_1(k4_waybel34(sK6)) | ~ v2_altcat_1(k4_waybel34(sK6)) | v3_struct_0(k4_waybel34(sK6))) | (spl38_14 | ~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18)), inference(subsumption_resolution, [], [f3034, f1943])).
fof(f3034, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | ~ l2_altcat_1(k4_waybel34(sK6)) | ~ v12_altcat_1(k4_waybel34(sK6)) | ~ v11_altcat_1(k4_waybel34(sK6)) | ~ v2_altcat_1(k4_waybel34(sK6)) | v3_struct_0(k4_waybel34(sK6))) | (~ spl38_15 | ~ spl38_16 | ~ spl38_17 | ~ spl38_18)), inference(subsumption_resolution, [], [f3033, f1947])).
fof(f3033, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | ~ l2_altcat_1(k4_waybel34(sK6)) | ~ v12_altcat_1(k4_waybel34(sK6)) | ~ v11_altcat_1(k4_waybel34(sK6)) | ~ v2_altcat_1(k4_waybel34(sK6)) | v3_struct_0(k4_waybel34(sK6))) | (~ spl38_16 | ~ spl38_17 | ~ spl38_18)), inference(subsumption_resolution, [], [f3032, f1951])).
fof(f3032, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v11_altcat_1(k5_waybel34(sK6)) | ~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | ~ l2_altcat_1(k4_waybel34(sK6)) | ~ v12_altcat_1(k4_waybel34(sK6)) | ~ v11_altcat_1(k4_waybel34(sK6)) | ~ v2_altcat_1(k4_waybel34(sK6)) | v3_struct_0(k4_waybel34(sK6))) | (~ spl38_17 | ~ spl38_18)), inference(subsumption_resolution, [], [f3031, f1955])).
fof(f3031, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v12_altcat_1(k5_waybel34(sK6)) | ~ v11_altcat_1(k5_waybel34(sK6)) | ~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | ~ l2_altcat_1(k4_waybel34(sK6)) | ~ v12_altcat_1(k4_waybel34(sK6)) | ~ v11_altcat_1(k4_waybel34(sK6)) | ~ v2_altcat_1(k4_waybel34(sK6)) | v3_struct_0(k4_waybel34(sK6))) | ~ spl38_18), inference(subsumption_resolution, [], [f2262, f1959])).
fof(f2262, plain, ! [X0, X1] : (r3_yellow20(k5_waybel34(sK6), k4_waybel34(sK6), k7_waybel34(sK6), X0, X1) | ~ r3_yellow20(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK6)) | ~ v3_altcat_2(X0, k5_waybel34(sK6)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK6)) | ~ v3_altcat_2(X1, k4_waybel34(sK6)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ m2_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | ~ l2_altcat_1(k5_waybel34(sK6)) | ~ v12_altcat_1(k5_waybel34(sK6)) | ~ v11_altcat_1(k5_waybel34(sK6)) | ~ v2_altcat_1(k5_waybel34(sK6)) | v3_struct_0(k5_waybel34(sK6)) | ~ l2_altcat_1(k4_waybel34(sK6)) | ~ v12_altcat_1(k4_waybel34(sK6)) | ~ v11_altcat_1(k4_waybel34(sK6)) | ~ v2_altcat_1(k4_waybel34(sK6)) | v3_struct_0(k4_waybel34(sK6))), inference(superposition, [], [f736, f1003])).
fof(f1003, plain, (k7_waybel34(sK6) = k15_functor0(k4_waybel34(sK6), k5_waybel34(sK6), k6_waybel34(sK6))), inference(resolution, [], [f726, f412])).
fof(f726, plain, ! [X0] : (v2_setfam_1(X0) | (k7_waybel34(X0) = k15_functor0(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0)))), inference(cnf_transformation, [], [f311])).
fof(f311, plain, ! [X0] : (((k6_waybel34(X0) = k15_functor0(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0))) & (k7_waybel34(X0) = k15_functor0(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0)))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f122])).
fof(f122, plain, ! [X0] : (~ v2_setfam_1(X0) => ((k6_waybel34(X0) = k15_functor0(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0))) & (k7_waybel34(X0) = k15_functor0(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', t18_waybel34)).
fof(f736, plain, ! [X4, X2, X0, X3, X1] : (r3_yellow20(X1, X0, k15_functor0(X0, X1, X2), X4, X3) | ~ r3_yellow20(X0, X1, X2, X3, X4) | ~ m1_altcat_2(X4, X1) | ~ v3_altcat_2(X4, X1) | ~ v2_altcat_1(X4) | v3_struct_0(X4) | ~ m1_altcat_2(X3, X0) | ~ v3_altcat_2(X3, X0) | ~ v2_altcat_1(X3) | v3_struct_0(X3) | ~ v21_functor0(X2, X0, X1) | ~ m2_functor0(X2, X0, X1) | ~ v16_functor0(X2, X0, X1) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ l2_altcat_1(X0) | ~ v12_altcat_1(X0) | ~ v11_altcat_1(X0) | ~ v2_altcat_1(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f321])).
fof(f321, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (r3_yellow20(X1, X0, k15_functor0(X0, X1, X2), X4, X3) | ~ r3_yellow20(X0, X1, X2, X3, X4) | ~ m1_altcat_2(X4, X1) | ~ v3_altcat_2(X4, X1) | ~ v2_altcat_1(X4) | v3_struct_0(X4)) | ~ m1_altcat_2(X3, X0) | ~ v3_altcat_2(X3, X0) | ~ v2_altcat_1(X3) | v3_struct_0(X3)) | ~ v21_functor0(X2, X0, X1) | ~ m2_functor0(X2, X0, X1) | ~ v16_functor0(X2, X0, X1)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1)) | ~ l2_altcat_1(X0) | ~ v12_altcat_1(X0) | ~ v11_altcat_1(X0) | ~ v2_altcat_1(X0) | v3_struct_0(X0)), inference(flattening, [], [f320])).
fof(f320, plain, ! [X0] : (! [X1] : (! [X2] : ((! [X3] : (! [X4] : ((r3_yellow20(X1, X0, k15_functor0(X0, X1, X2), X4, X3) | ~ r3_yellow20(X0, X1, X2, X3, X4)) | (~ m1_altcat_2(X4, X1) | ~ v3_altcat_2(X4, X1) | ~ v2_altcat_1(X4) | v3_struct_0(X4))) | (~ m1_altcat_2(X3, X0) | ~ v3_altcat_2(X3, X0) | ~ v2_altcat_1(X3) | v3_struct_0(X3))) | ~ v21_functor0(X2, X0, X1)) | (~ m2_functor0(X2, X0, X1) | ~ v16_functor0(X2, X0, X1))) | (~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1))) | (~ l2_altcat_1(X0) | ~ v12_altcat_1(X0) | ~ v11_altcat_1(X0) | ~ v2_altcat_1(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f128])).
fof(f128, plain, ! [X0] : ((l2_altcat_1(X0) & v12_altcat_1(X0) & v11_altcat_1(X0) & v2_altcat_1(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((l2_altcat_1(X1) & v12_altcat_1(X1) & v11_altcat_1(X1) & v2_altcat_1(X1) & ~ v3_struct_0(X1)) => ! [X2] : ((m2_functor0(X2, X0, X1) & v16_functor0(X2, X0, X1)) => (v21_functor0(X2, X0, X1) => ! [X3] : ((m1_altcat_2(X3, X0) & v3_altcat_2(X3, X0) & v2_altcat_1(X3) & ~ v3_struct_0(X3)) => ! [X4] : ((m1_altcat_2(X4, X1) & v3_altcat_2(X4, X1) & v2_altcat_1(X4) & ~ v3_struct_0(X4)) => (r3_yellow20(X0, X1, X2, X3, X4) => r3_yellow20(X1, X0, k15_functor0(X0, X1, X2), X4, X3)))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', t52_yellow20)).
fof(f3030, plain, spl38_19, inference(avatar_contradiction_clause, [], [f3029])).
fof(f3029, plain, ($false | spl38_19), inference(subsumption_resolution, [], [f3028, f412])).
fof(f3028, plain, (v2_setfam_1(sK6) | spl38_19), inference(resolution, [], [f1964, f552])).
fof(f552, plain, ! [X0] : (v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f245])).
fof(f1964, plain, (~ v16_functor0(k6_waybel34(sK6), k4_waybel34(sK6), k5_waybel34(sK6)) | spl38_19), inference(avatar_component_clause, [], [f1962])).
fof(f2275, plain, ~ spl38_14, inference(avatar_contradiction_clause, [], [f2274])).
fof(f2274, plain, ($false | ~ spl38_14), inference(subsumption_resolution, [], [f2273, f412])).
fof(f2273, plain, (v2_setfam_1(sK6) | ~ spl38_14), inference(resolution, [], [f2271, f644])).
fof(f644, plain, ! [X0] : (sP5(X0) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f338])).
fof(f338, plain, ! [X0] : (sP5(X0) | v2_setfam_1(X0)), inference(definition_folding, [], [f287, e337])).
fof(f337, plain, ! [X0] : ((v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0))) | ~ sP5(X0)), inference(usedef, [], [e337])).
fof(e337, plain, ! [X0] : (sP5(X0) <=> (v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f287, plain, ! [X0] : ((v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f176])).
fof(f176, plain, ! [X0] : (~ v2_setfam_1(X0) => (v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0)))), inference(pure_predicate_removal, [], [f96])).
fof(f96, plain, ! [X0] : (~ v2_setfam_1(X0) => (v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v4_yellow18(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', fc6_waybel34)).
fof(f2271, plain, (~ sP5(sK6) | ~ spl38_14), inference(resolution, [], [f1944, f632])).
fof(f632, plain, ! [X0] : (~ v3_struct_0(k5_waybel34(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f374])).
fof(f374, plain, ! [X0] : ((v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0))) | ~ sP5(X0)), inference(nnf_transformation, [], [f337])).
fof(f1944, plain, (v3_struct_0(k5_waybel34(sK6)) | ~ spl38_14), inference(avatar_component_clause, [], [f1942])).
fof(f2252, plain, spl38_18, inference(avatar_contradiction_clause, [], [f2251])).
fof(f2251, plain, ($false | spl38_18), inference(subsumption_resolution, [], [f2225, f762])).
fof(f762, plain, v2_setfam_1(k1_xboole_0), inference(resolution, [], [f466, f626])).
fof(f626, plain, v1_xboole_0(k1_xboole_0), inference(cnf_transformation, [], [f95])).
fof(f95, plain, (v5_membered(k1_xboole_0) & v4_membered(k1_xboole_0) & v3_membered(k1_xboole_0) & v2_membered(k1_xboole_0) & v1_membered(k1_xboole_0) & v1_xboole_0(k1_xboole_0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', fc6_membered)).
fof(f2225, plain, (~ v2_setfam_1(k1_xboole_0) | spl38_18), inference(backward_demodulation, [], [f412, f2220])).
fof(f2220, plain, ((k1_xboole_0 = sK6) | spl38_18), inference(resolution, [], [f2211, f739])).
fof(f739, plain, ! [X0] : (~ v1_xboole_0(X0) | (k1_xboole_0 = X0)), inference(cnf_transformation, [], [f324])).
fof(f324, plain, ! [X0] : ((k1_xboole_0 = X0) | ~ v1_xboole_0(X0)), inference(ennf_transformation, [], [f131])).
fof(f131, plain, ! [X0] : (v1_xboole_0(X0) => (k1_xboole_0 = X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', t6_boole)).
fof(f2211, plain, (v1_xboole_0(sK6) | spl38_18), inference(resolution, [], [f1960, f550])).
fof(f550, plain, ! [X0] : (l2_altcat_1(k5_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f244])).
fof(f244, plain, ! [X0] : ((l2_altcat_1(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0))) | v1_xboole_0(X0)), inference(ennf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : (~ v1_xboole_0(X0) => (l2_altcat_1(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', dt_k5_waybel34)).
fof(f1960, plain, (~ l2_altcat_1(k5_waybel34(sK6)) | spl38_18), inference(avatar_component_clause, [], [f1958])).
fof(f2209, plain, spl38_17, inference(avatar_contradiction_clause, [], [f2208])).
fof(f2208, plain, ($false | spl38_17), inference(subsumption_resolution, [], [f2207, f412])).
fof(f2207, plain, (v2_setfam_1(sK6) | spl38_17), inference(resolution, [], [f2205, f644])).
fof(f2205, plain, (~ sP5(sK6) | spl38_17), inference(resolution, [], [f1956, f637])).
fof(f637, plain, ! [X0] : (v12_altcat_1(k5_waybel34(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f374])).
fof(f1956, plain, (~ v12_altcat_1(k5_waybel34(sK6)) | spl38_17), inference(avatar_component_clause, [], [f1954])).
fof(f2201, plain, spl38_16, inference(avatar_contradiction_clause, [], [f2200])).
fof(f2200, plain, ($false | spl38_16), inference(subsumption_resolution, [], [f2199, f412])).
fof(f2199, plain, (v2_setfam_1(sK6) | spl38_16), inference(resolution, [], [f2194, f644])).
fof(f2194, plain, (~ sP5(sK6) | spl38_16), inference(resolution, [], [f1952, f636])).
fof(f636, plain, ! [X0] : (v11_altcat_1(k5_waybel34(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f374])).
fof(f1952, plain, (~ v11_altcat_1(k5_waybel34(sK6)) | spl38_16), inference(avatar_component_clause, [], [f1950])).
fof(f2175, plain, spl38_15, inference(avatar_contradiction_clause, [], [f2174])).
fof(f2174, plain, ($false | spl38_15), inference(subsumption_resolution, [], [f2173, f412])).
fof(f2173, plain, (v2_setfam_1(sK6) | spl38_15), inference(resolution, [], [f2168, f644])).
fof(f2168, plain, (~ sP5(sK6) | spl38_15), inference(resolution, [], [f1948, f633])).
fof(f633, plain, ! [X0] : (v2_altcat_1(k5_waybel34(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f374])).
fof(f1948, plain, (~ v2_altcat_1(k5_waybel34(sK6)) | spl38_15), inference(avatar_component_clause, [], [f1946])).
fof(f2109, plain, ~ spl38_9, inference(avatar_contradiction_clause, [], [f2108])).
fof(f2108, plain, ($false | ~ spl38_9), inference(subsumption_resolution, [], [f2107, f412])).
fof(f2107, plain, (v2_setfam_1(sK6) | ~ spl38_9), inference(resolution, [], [f2105, f625])).
fof(f625, plain, ! [X0] : (sP4(X0) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f336])).
fof(f336, plain, ! [X0] : (sP4(X0) | v2_setfam_1(X0)), inference(definition_folding, [], [f286, e335])).
fof(f335, plain, ! [X0] : ((v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0))) | ~ sP4(X0)), inference(usedef, [], [e335])).
fof(e335, plain, ! [X0] : (sP4(X0) <=> (v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f286, plain, ! [X0] : ((v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f172])).
fof(f172, plain, ! [X0] : (~ v2_setfam_1(X0) => (v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0)))), inference(pure_predicate_removal, [], [f94])).
fof(f94, plain, ! [X0] : (~ v2_setfam_1(X0) => (v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v4_yellow18(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT377+1.p', fc5_waybel34)).
fof(f2105, plain, (~ sP4(sK6) | ~ spl38_9), inference(resolution, [], [f1924, f613])).
fof(f613, plain, ! [X0] : (~ v3_struct_0(k4_waybel34(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f373])).
fof(f373, plain, ! [X0] : ((v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0))) | ~ sP4(X0)), inference(nnf_transformation, [], [f335])).
fof(f1924, plain, (v3_struct_0(k4_waybel34(sK6)) | ~ spl38_9), inference(avatar_component_clause, [], [f1922])).
fof(f2095, plain, spl38_13, inference(avatar_contradiction_clause, [], [f2094])).
fof(f2094, plain, ($false | spl38_13), inference(subsumption_resolution, [], [f2077, f762])).
fof(f2077, plain, (~ v2_setfam_1(k1_xboole_0) | spl38_13), inference(backward_demodulation, [], [f412, f2072])).
fof(f2072, plain, ((k1_xboole_0 = sK6) | spl38_13), inference(resolution, [], [f2060, f739])).
fof(f2060, plain, (v1_xboole_0(sK6) | spl38_13), inference(resolution, [], [f1940, f543])).
fof(f1940, plain, (~ l2_altcat_1(k4_waybel34(sK6)) | spl38_13), inference(avatar_component_clause, [], [f1938])).
fof(f2052, plain, spl38_12, inference(avatar_contradiction_clause, [], [f2051])).
fof(f2051, plain, ($false | spl38_12), inference(subsumption_resolution, [], [f2050, f412])).
fof(f2050, plain, (v2_setfam_1(sK6) | spl38_12), inference(resolution, [], [f2045, f625])).
fof(f2045, plain, (~ sP4(sK6) | spl38_12), inference(resolution, [], [f1936, f618])).
fof(f618, plain, ! [X0] : (v12_altcat_1(k4_waybel34(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f373])).
fof(f1936, plain, (~ v12_altcat_1(k4_waybel34(sK6)) | spl38_12), inference(avatar_component_clause, [], [f1934])).
fof(f2033, plain, spl38_10, inference(avatar_contradiction_clause, [], [f2032])).
fof(f2032, plain, ($false | spl38_10), inference(subsumption_resolution, [], [f2031, f412])).
fof(f2031, plain, (v2_setfam_1(sK6) | spl38_10), inference(resolution, [], [f2026, f625])).
fof(f2026, plain, (~ sP4(sK6) | spl38_10), inference(resolution, [], [f1928, f614])).
fof(f614, plain, ! [X0] : (v2_altcat_1(k4_waybel34(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f373])).
fof(f1928, plain, (~ v2_altcat_1(k4_waybel34(sK6)) | spl38_10), inference(avatar_component_clause, [], [f1926])).
fof(f2006, plain, spl38_11, inference(avatar_contradiction_clause, [], [f2005])).
fof(f2005, plain, ($false | spl38_11), inference(subsumption_resolution, [], [f2004, f412])).
fof(f2004, plain, (v2_setfam_1(sK6) | spl38_11), inference(resolution, [], [f1999, f625])).
fof(f1999, plain, (~ sP4(sK6) | spl38_11), inference(resolution, [], [f1932, f617])).
fof(f617, plain, ! [X0] : (v11_altcat_1(k4_waybel34(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f373])).
fof(f1932, plain, (~ v11_altcat_1(k4_waybel34(sK6)) | spl38_11), inference(avatar_component_clause, [], [f1930])).