fof(f410207, plain, $false, inference(avatar_sat_refutation, [], [f1286, f1291, f1311, f1316, f1321, f1326, f1334, f1335, f1753, f3468, f3687, f4024, f4090, f4519, f4615, f4631, f8130, f8149, f8168, f8246, f17172, f24575, f31095, f31132, f31189, f32302, f35175, f36138, f39742, f51196, f51778, f119268, f265972, f371800, f371812, f371848, f371905, f371907, f371965, f371967, f410202])).
fof(f410202, plain, (spl60_2 | ~ spl60_3 | ~ spl60_7 | ~ spl60_8 | ~ spl60_9 | ~ spl60_10 | ~ spl60_11 | ~ spl60_12 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_2156 | ~ spl60_31634 | ~ spl60_31636), inference(avatar_contradiction_clause, [], [f410201])).
fof(f410201, plain, ($false | (spl60_2 | ~ spl60_3 | ~ spl60_7 | ~ spl60_8 | ~ spl60_9 | ~ spl60_10 | ~ spl60_11 | ~ spl60_12 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_2156 | ~ spl60_31634 | ~ spl60_31636)), inference(subsumption_resolution, [], [f410200, f371820])).
fof(f371820, plain, (~ v1_t_0topsp(sK15, sK17, sK16) | (spl60_2 | ~ spl60_3)), inference(forward_demodulation, [], [f1285, f1290])).
fof(f1290, plain, ((sK15 = sK18) | ~ spl60_3), inference(avatar_component_clause, [], [f1288])).
fof(f1288, plain, (spl60_3 <=> (sK15 = sK18)), introduced(avatar_definition, [new_symbols(naming, [spl60_3])])).
fof(f1285, plain, (~ v1_t_0topsp(sK18, sK17, sK16) | spl60_2), inference(avatar_component_clause, [], [f1283])).
fof(f1283, plain, (spl60_2 <=> v1_t_0topsp(sK18, sK17, sK16)), introduced(avatar_definition, [new_symbols(naming, [spl60_2])])).
fof(f410200, plain, (v1_t_0topsp(sK15, sK17, sK16) | (~ spl60_7 | ~ spl60_8 | ~ spl60_9 | ~ spl60_10 | ~ spl60_11 | ~ spl60_12 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_2156 | ~ spl60_31634 | ~ spl60_31636)), inference(subsumption_resolution, [], [f410199, f374900])).
fof(f374900, plain, (v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK16)) | (~ spl60_11 | spl60_26 | ~ spl60_67 | ~ spl60_2156 | ~ spl60_31636)), inference(backward_demodulation, [], [f4542, f374897])).
fof(f374897, plain, ((u1_struct_0(sK16) = u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_67 | ~ spl60_2156 | ~ spl60_31636)), inference(backward_demodulation, [], [f266025, f374877])).
fof(f374877, plain, ((u1_struct_0(sK16) = u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | (~ spl60_2156 | ~ spl60_31636)), inference(trivial_inequality_removal, [], [f374872])).
fof(f374872, plain, (~ (k2_yellow_2(sK14, sK14, sK15) = k2_yellow_2(sK14, sK14, sK15)) | (u1_struct_0(sK16) = u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | (~ spl60_2156 | ~ spl60_31636)), inference(superposition, [], [f35174, f371964])).
fof(f371964, plain, ((k2_yellow_2(sK14, sK14, sK15) = g1_orders_2(u1_struct_0(sK16), u1_orders_2(sK16))) | ~ spl60_31636), inference(avatar_component_clause, [], [f371962])).
fof(f371962, plain, (spl60_31636 <=> (k2_yellow_2(sK14, sK14, sK15) = g1_orders_2(u1_struct_0(sK16), u1_orders_2(sK16)))), introduced(avatar_definition, [new_symbols(naming, [spl60_31636])])).
fof(f35174, plain, (! [X2, X3] : (~ (g1_orders_2(X2, X3) = k2_yellow_2(sK14, sK14, sK15)) | (u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))) = X2)) | ~ spl60_2156), inference(avatar_component_clause, [], [f35173])).
fof(f35173, plain, (spl60_2156 <=> ! [X3, X2] : (~ (g1_orders_2(X2, X3) = k2_yellow_2(sK14, sK14, sK15)) | (u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))) = X2))), introduced(avatar_definition, [new_symbols(naming, [spl60_2156])])).
fof(f266025, plain, ((u1_struct_0(k2_yellow_2(sK14, sK14, sK15)) = u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | (~ spl60_67 | ~ spl60_2156)), inference(trivial_inequality_removal, [], [f265973])).
fof(f265973, plain, (~ (k2_yellow_2(sK14, sK14, sK15) = k2_yellow_2(sK14, sK14, sK15)) | (u1_struct_0(k2_yellow_2(sK14, sK14, sK15)) = u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | (~ spl60_67 | ~ spl60_2156)), inference(superposition, [], [f35174, f3467])).
fof(f3467, plain, ((k2_yellow_2(sK14, sK14, sK15) = g1_orders_2(u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_orders_2(k2_yellow_2(sK14, sK14, sK15)))) | ~ spl60_67), inference(avatar_component_clause, [], [f3465])).
fof(f3465, plain, (spl60_67 <=> (k2_yellow_2(sK14, sK14, sK15) = g1_orders_2(u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_orders_2(k2_yellow_2(sK14, sK14, sK15))))), introduced(avatar_definition, [new_symbols(naming, [spl60_67])])).
fof(f4542, plain, (v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4541, f1653])).
fof(f1653, plain, (~ v3_struct_0(sK14) | spl60_26), inference(avatar_component_clause, [], [f1652])).
fof(f1652, plain, (spl60_26 <=> v3_struct_0(sK14)), introduced(avatar_definition, [new_symbols(naming, [spl60_26])])).
fof(f4541, plain, (v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4540, f634])).
fof(f634, plain, l1_orders_2(sK14), inference(cnf_transformation, [], [f524])).
fof(f524, plain, ((((((~ v1_t_0topsp(sK18, sK17, sK16) & (sK15 = sK18) & m2_relset_1(sK18, u1_struct_0(sK17), u1_struct_0(sK16)) & v1_funct_2(sK18, u1_struct_0(sK17), u1_struct_0(sK16)) & v1_funct_1(sK18)) & m1_yellow_9(sK17, sK14) & v4_waybel11(sK17)) & m1_yellow_9(sK16, k2_yellow_2(sK14, sK14, sK15)) & v4_waybel11(sK16)) | ~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)) & (! [X5] : (! [X6] : (! [X7] : (v1_t_0topsp(X7, X6, X5) | ~ (sK15 = X7) | ~ m2_relset_1(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_2(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_1(X7)) | ~ m1_yellow_9(X6, sK14) | ~ v4_waybel11(X6)) | ~ m1_yellow_9(X5, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X5)) | v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)) & m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) & v7_waybel_1(sK15, sK14) & v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) & v1_funct_1(sK15)) & l1_orders_2(sK14) & v3_lattice3(sK14) & v2_lattice3(sK14) & v1_lattice3(sK14) & v4_orders_2(sK14) & v3_orders_2(sK14) & v2_orders_2(sK14)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14, sK15, sK16, sK17, sK18])], [f518, f523, f522, f521, f520, f519])).
fof(f519, plain, (? [X0] : (? [X1] : ((? [X2] : (? [X3] : (? [X4] : (~ v1_t_0topsp(X4, X3, X2) & (X1 = X4) & m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_1(X4)) & m1_yellow_9(X3, X0) & v4_waybel11(X3)) & m1_yellow_9(X2, k2_yellow_2(X0, X0, X1)) & v4_waybel11(X2)) | ~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0)) & (! [X5] : (! [X6] : (! [X7] : (v1_t_0topsp(X7, X6, X5) | ~ (X1 = X7) | ~ m2_relset_1(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_2(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_1(X7)) | ~ m1_yellow_9(X6, X0) | ~ v4_waybel11(X6)) | ~ m1_yellow_9(X5, k2_yellow_2(X0, X0, X1)) | ~ v4_waybel11(X5)) | v4_waybel_0(k2_yellow_2(X0, X0, X1), X0)) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (? [X1] : ((? [X2] : (? [X3] : (? [X4] : (~ v1_t_0topsp(X4, X3, X2) & (X1 = X4) & m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_1(X4)) & m1_yellow_9(X3, sK14) & v4_waybel11(X3)) & m1_yellow_9(X2, k2_yellow_2(sK14, sK14, X1)) & v4_waybel11(X2)) | ~ v4_waybel_0(k2_yellow_2(sK14, sK14, X1), sK14)) & (! [X5] : (! [X6] : (! [X7] : (v1_t_0topsp(X7, X6, X5) | ~ (X1 = X7) | ~ m2_relset_1(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_2(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_1(X7)) | ~ m1_yellow_9(X6, sK14) | ~ v4_waybel11(X6)) | ~ m1_yellow_9(X5, k2_yellow_2(sK14, sK14, X1)) | ~ v4_waybel11(X5)) | v4_waybel_0(k2_yellow_2(sK14, sK14, X1), sK14)) & m2_relset_1(X1, u1_struct_0(sK14), u1_struct_0(sK14)) & v7_waybel_1(X1, sK14) & v1_funct_2(X1, u1_struct_0(sK14), u1_struct_0(sK14)) & v1_funct_1(X1)) & l1_orders_2(sK14) & v3_lattice3(sK14) & v2_lattice3(sK14) & v1_lattice3(sK14) & v4_orders_2(sK14) & v3_orders_2(sK14) & v2_orders_2(sK14))), introduced(choice_axiom, [])).
fof(f520, plain, (? [X1] : ((? [X2] : (? [X3] : (? [X4] : (~ v1_t_0topsp(X4, X3, X2) & (X1 = X4) & m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_1(X4)) & m1_yellow_9(X3, sK14) & v4_waybel11(X3)) & m1_yellow_9(X2, k2_yellow_2(sK14, sK14, X1)) & v4_waybel11(X2)) | ~ v4_waybel_0(k2_yellow_2(sK14, sK14, X1), sK14)) & (! [X5] : (! [X6] : (! [X7] : (v1_t_0topsp(X7, X6, X5) | ~ (X1 = X7) | ~ m2_relset_1(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_2(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_1(X7)) | ~ m1_yellow_9(X6, sK14) | ~ v4_waybel11(X6)) | ~ m1_yellow_9(X5, k2_yellow_2(sK14, sK14, X1)) | ~ v4_waybel11(X5)) | v4_waybel_0(k2_yellow_2(sK14, sK14, X1), sK14)) & m2_relset_1(X1, u1_struct_0(sK14), u1_struct_0(sK14)) & v7_waybel_1(X1, sK14) & v1_funct_2(X1, u1_struct_0(sK14), u1_struct_0(sK14)) & v1_funct_1(X1)) => ((? [X2] : (? [X3] : (? [X4] : (~ v1_t_0topsp(X4, X3, X2) & (sK15 = X4) & m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_1(X4)) & m1_yellow_9(X3, sK14) & v4_waybel11(X3)) & m1_yellow_9(X2, k2_yellow_2(sK14, sK14, sK15)) & v4_waybel11(X2)) | ~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)) & (! [X5] : (! [X6] : (! [X7] : (v1_t_0topsp(X7, X6, X5) | ~ (sK15 = X7) | ~ m2_relset_1(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_2(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_1(X7)) | ~ m1_yellow_9(X6, sK14) | ~ v4_waybel11(X6)) | ~ m1_yellow_9(X5, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X5)) | v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)) & m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) & v7_waybel_1(sK15, sK14) & v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) & v1_funct_1(sK15))), introduced(choice_axiom, [])).
fof(f521, plain, (? [X2] : (? [X3] : (? [X4] : (~ v1_t_0topsp(X4, X3, X2) & (sK15 = X4) & m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_1(X4)) & m1_yellow_9(X3, sK14) & v4_waybel11(X3)) & m1_yellow_9(X2, k2_yellow_2(sK14, sK14, sK15)) & v4_waybel11(X2)) => (? [X3] : (? [X4] : (~ v1_t_0topsp(X4, X3, sK16) & (sK15 = X4) & m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(sK16)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(sK16)) & v1_funct_1(X4)) & m1_yellow_9(X3, sK14) & v4_waybel11(X3)) & m1_yellow_9(sK16, k2_yellow_2(sK14, sK14, sK15)) & v4_waybel11(sK16))), introduced(choice_axiom, [])).
fof(f522, plain, (? [X3] : (? [X4] : (~ v1_t_0topsp(X4, X3, sK16) & (sK15 = X4) & m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(sK16)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(sK16)) & v1_funct_1(X4)) & m1_yellow_9(X3, sK14) & v4_waybel11(X3)) => (? [X4] : (~ v1_t_0topsp(X4, sK17, sK16) & (sK15 = X4) & m2_relset_1(X4, u1_struct_0(sK17), u1_struct_0(sK16)) & v1_funct_2(X4, u1_struct_0(sK17), u1_struct_0(sK16)) & v1_funct_1(X4)) & m1_yellow_9(sK17, sK14) & v4_waybel11(sK17))), introduced(choice_axiom, [])).
fof(f523, plain, (? [X4] : (~ v1_t_0topsp(X4, sK17, sK16) & (sK15 = X4) & m2_relset_1(X4, u1_struct_0(sK17), u1_struct_0(sK16)) & v1_funct_2(X4, u1_struct_0(sK17), u1_struct_0(sK16)) & v1_funct_1(X4)) => (~ v1_t_0topsp(sK18, sK17, sK16) & (sK15 = sK18) & m2_relset_1(sK18, u1_struct_0(sK17), u1_struct_0(sK16)) & v1_funct_2(sK18, u1_struct_0(sK17), u1_struct_0(sK16)) & v1_funct_1(sK18))), introduced(choice_axiom, [])).
fof(f518, plain, ? [X0] : (? [X1] : ((? [X2] : (? [X3] : (? [X4] : (~ v1_t_0topsp(X4, X3, X2) & (X1 = X4) & m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_1(X4)) & m1_yellow_9(X3, X0) & v4_waybel11(X3)) & m1_yellow_9(X2, k2_yellow_2(X0, X0, X1)) & v4_waybel11(X2)) | ~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0)) & (! [X5] : (! [X6] : (! [X7] : (v1_t_0topsp(X7, X6, X5) | ~ (X1 = X7) | ~ m2_relset_1(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_2(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_1(X7)) | ~ m1_yellow_9(X6, X0) | ~ v4_waybel11(X6)) | ~ m1_yellow_9(X5, k2_yellow_2(X0, X0, X1)) | ~ v4_waybel11(X5)) | v4_waybel_0(k2_yellow_2(X0, X0, X1), X0)) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(rectify, [], [f517])).
fof(f517, plain, ? [X0] : (? [X1] : ((? [X2] : (? [X3] : (? [X4] : (~ v1_t_0topsp(X4, X3, X2) & (X1 = X4) & m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_1(X4)) & m1_yellow_9(X3, X0) & v4_waybel11(X3)) & m1_yellow_9(X2, k2_yellow_2(X0, X0, X1)) & v4_waybel11(X2)) | ~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0)) & (! [X2] : (! [X3] : (! [X4] : (v1_t_0topsp(X4, X3, X2) | ~ (X1 = X4) | ~ m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) | ~ v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) | ~ v1_funct_1(X4)) | ~ m1_yellow_9(X3, X0) | ~ v4_waybel11(X3)) | ~ m1_yellow_9(X2, k2_yellow_2(X0, X0, X1)) | ~ v4_waybel11(X2)) | v4_waybel_0(k2_yellow_2(X0, X0, X1), X0)) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(flattening, [], [f516])).
fof(f516, plain, ? [X0] : (? [X1] : (((? [X2] : (? [X3] : (? [X4] : (~ v1_t_0topsp(X4, X3, X2) & (X1 = X4) & m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_1(X4)) & m1_yellow_9(X3, X0) & v4_waybel11(X3)) & m1_yellow_9(X2, k2_yellow_2(X0, X0, X1)) & v4_waybel11(X2)) | ~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0)) & (! [X2] : (! [X3] : (! [X4] : (v1_t_0topsp(X4, X3, X2) | ~ (X1 = X4) | ~ m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) | ~ v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) | ~ v1_funct_1(X4)) | ~ m1_yellow_9(X3, X0) | ~ v4_waybel11(X3)) | ~ m1_yellow_9(X2, k2_yellow_2(X0, X0, X1)) | ~ v4_waybel11(X2)) | v4_waybel_0(k2_yellow_2(X0, X0, X1), X0))) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(nnf_transformation, [], [f244])).
fof(f244, plain, ? [X0] : (? [X1] : (~ (v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> ! [X2] : (! [X3] : (! [X4] : (v1_t_0topsp(X4, X3, X2) | ~ (X1 = X4) | ~ m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) | ~ v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) | ~ v1_funct_1(X4)) | ~ m1_yellow_9(X3, X0) | ~ v4_waybel11(X3)) | ~ m1_yellow_9(X2, k2_yellow_2(X0, X0, X1)) | ~ v4_waybel11(X2))) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(flattening, [], [f243])).
fof(f243, plain, ? [X0] : (? [X1] : (~ (v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> ! [X2] : (! [X3] : (! [X4] : ((v1_t_0topsp(X4, X3, X2) | ~ (X1 = X4)) | (~ m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) | ~ v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) | ~ v1_funct_1(X4))) | (~ m1_yellow_9(X3, X0) | ~ v4_waybel11(X3))) | (~ m1_yellow_9(X2, k2_yellow_2(X0, X0, X1)) | ~ v4_waybel11(X2)))) & (m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1))) & (l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> ! [X2] : ((m1_yellow_9(X2, k2_yellow_2(X0, X0, X1)) & v4_waybel11(X2)) => ! [X3] : ((m1_yellow_9(X3, X0) & v4_waybel11(X3)) => ! [X4] : ((m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_1(X4)) => ((X1 = X4) => v1_t_0topsp(X4, X3, X2)))))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> ! [X2] : ((m1_yellow_9(X2, k2_yellow_2(X0, X0, X1)) & v4_waybel11(X2)) => ! [X3] : ((m1_yellow_9(X3, X0) & v4_waybel11(X3)) => ! [X4] : ((m2_relset_1(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_2(X4, u1_struct_0(X3), u1_struct_0(X2)) & v1_funct_1(X4)) => ((X1 = X4) => v1_t_0topsp(X4, X3, X2)))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', t41_waybel34)).
fof(f4540, plain, (v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4539, f1329])).
fof(f1329, plain, (v1_funct_1(sK15) | ~ spl60_11), inference(avatar_component_clause, [], [f1328])).
fof(f1328, plain, (spl60_11 <=> v1_funct_1(sK15)), introduced(avatar_definition, [new_symbols(naming, [spl60_11])])).
fof(f4539, plain, (v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4538, f636])).
fof(f636, plain, v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)), inference(cnf_transformation, [], [f524])).
fof(f4538, plain, (v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4537, f1864])).
fof(f1864, plain, m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)), inference(resolution, [], [f1255, f638])).
fof(f638, plain, m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)), inference(cnf_transformation, [], [f524])).
fof(f1255, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f625])).
fof(f625, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f164])).
fof(f164, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', redefinition_m2_relset_1)).
fof(f4537, plain, (v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(duplicate_literal_removal, [], [f4536])).
fof(f4536, plain, (v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(superposition, [], [f859, f3872])).
fof(f3872, plain, ((sK15 = k2_waybel_1(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3871, f1653])).
fof(f3871, plain, ((sK15 = k2_waybel_1(sK14, sK14, sK15)) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3870, f634])).
fof(f3870, plain, ((sK15 = k2_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3869, f1329])).
fof(f3869, plain, ((sK15 = k2_waybel_1(sK14, sK14, sK15)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f3863, f636])).
fof(f3863, plain, ((sK15 = k2_waybel_1(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(duplicate_literal_removal, [], [f3860])).
fof(f3860, plain, ((sK15 = k2_waybel_1(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(resolution, [], [f1262, f638])).
fof(f1262, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | (k2_waybel_1(X0, X1, X2) = X2) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f476])).
fof(f476, plain, ! [X0] : (! [X1] : (! [X2] : ((k2_waybel_1(X0, X1, X2) = X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | v3_struct_0(X1)) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f475])).
fof(f475, plain, ! [X0] : (! [X1] : (! [X2] : ((k2_waybel_1(X0, X1, X2) = X2) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | v3_struct_0(X1))) | (~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f169])).
fof(f169, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((l1_orders_2(X1) & ~ v3_struct_0(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => (k2_waybel_1(X0, X1, X2) = X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', t32_waybel_1)).
fof(f859, plain, ! [X2, X0, X1] : (v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f363])).
fof(f363, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f362])).
fof(f362, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f78])).
fof(f78, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', dt_k2_waybel_1)).
fof(f410199, plain, (~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK16)) | v1_t_0topsp(sK15, sK17, sK16) | (~ spl60_7 | ~ spl60_8 | ~ spl60_9 | ~ spl60_10 | ~ spl60_11 | ~ spl60_12 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_2156 | ~ spl60_31634 | ~ spl60_31636)), inference(subsumption_resolution, [], [f410198, f1310])).
fof(f1310, plain, (m1_yellow_9(sK17, sK14) | ~ spl60_7), inference(avatar_component_clause, [], [f1308])).
fof(f1308, plain, (spl60_7 <=> m1_yellow_9(sK17, sK14)), introduced(avatar_definition, [new_symbols(naming, [spl60_7])])).
fof(f410198, plain, (~ m1_yellow_9(sK17, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK16)) | v1_t_0topsp(sK15, sK17, sK16) | (~ spl60_8 | ~ spl60_9 | ~ spl60_10 | ~ spl60_11 | ~ spl60_12 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_2156 | ~ spl60_31634 | ~ spl60_31636)), inference(subsumption_resolution, [], [f410197, f1315])).
fof(f1315, plain, (v4_waybel11(sK17) | ~ spl60_8), inference(avatar_component_clause, [], [f1313])).
fof(f1313, plain, (spl60_8 <=> v4_waybel11(sK17)), introduced(avatar_definition, [new_symbols(naming, [spl60_8])])).
fof(f410197, plain, (~ v4_waybel11(sK17) | ~ m1_yellow_9(sK17, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK16)) | v1_t_0topsp(sK15, sK17, sK16) | (~ spl60_9 | ~ spl60_10 | ~ spl60_11 | ~ spl60_12 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_2156 | ~ spl60_31634 | ~ spl60_31636)), inference(subsumption_resolution, [], [f405319, f374901])).
fof(f374901, plain, (m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK16)) | (~ spl60_11 | spl60_26 | ~ spl60_67 | ~ spl60_2156 | ~ spl60_31636)), inference(backward_demodulation, [], [f4568, f374897])).
fof(f4568, plain, (m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4567, f1653])).
fof(f4567, plain, (m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4566, f634])).
fof(f4566, plain, (m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4565, f1329])).
fof(f4565, plain, (m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4564, f636])).
fof(f4564, plain, (m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4551, f1864])).
fof(f4551, plain, (m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(duplicate_literal_removal, [], [f4550])).
fof(f4550, plain, (m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | (~ spl60_11 | spl60_26)), inference(superposition, [], [f860, f3872])).
fof(f860, plain, ! [X2, X0, X1] : (m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f363])).
fof(f405319, plain, (~ m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK16)) | ~ v4_waybel11(sK17) | ~ m1_yellow_9(sK17, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK16)) | v1_t_0topsp(sK15, sK17, sK16) | (~ spl60_9 | ~ spl60_10 | ~ spl60_12 | ~ spl60_1818 | ~ spl60_31634)), inference(superposition, [], [f372051, f375469])).
fof(f375469, plain, ((u1_struct_0(sK14) = u1_struct_0(sK17)) | (~ spl60_1818 | ~ spl60_31634)), inference(trivial_inequality_removal, [], [f375466])).
fof(f375466, plain, (~ (g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14)) = g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14))) | (u1_struct_0(sK14) = u1_struct_0(sK17)) | (~ spl60_1818 | ~ spl60_31634)), inference(superposition, [], [f171034, f371904])).
fof(f371904, plain, ((g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14)) = g1_orders_2(u1_struct_0(sK17), u1_orders_2(sK17))) | ~ spl60_31634), inference(avatar_component_clause, [], [f371902])).
fof(f371902, plain, (spl60_31634 <=> (g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14)) = g1_orders_2(u1_struct_0(sK17), u1_orders_2(sK17)))), introduced(avatar_definition, [new_symbols(naming, [spl60_31634])])).
fof(f171034, plain, (! [X2, X3] : (~ (g1_orders_2(X2, X3) = g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14))) | (u1_struct_0(sK14) = X2)) | ~ spl60_1818), inference(backward_demodulation, [], [f31094, f170797])).
fof(f170797, plain, ((u1_struct_0(sK14) = u1_struct_0(sK26(sK14))) | ~ spl60_1818), inference(equality_resolution, [], [f31094])).
fof(f31094, plain, (! [X2, X3] : (~ (g1_orders_2(X2, X3) = g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14))) | (u1_struct_0(sK26(sK14)) = X2)) | ~ spl60_1818), inference(avatar_component_clause, [], [f31093])).
fof(f31093, plain, (spl60_1818 <=> ! [X3, X2] : (~ (g1_orders_2(X2, X3) = g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14))) | (u1_struct_0(sK26(sK14)) = X2))), introduced(avatar_definition, [new_symbols(naming, [spl60_1818])])).
fof(f372051, plain, (! [X0] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(sK16)) | ~ v4_waybel11(X0) | ~ m1_yellow_9(X0, sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(sK16)) | v1_t_0topsp(sK15, X0, sK16)) | (~ spl60_9 | ~ spl60_10 | ~ spl60_12)), inference(subsumption_resolution, [], [f372047, f1325])).
fof(f1325, plain, (v4_waybel11(sK16) | ~ spl60_10), inference(avatar_component_clause, [], [f1323])).
fof(f1323, plain, (spl60_10 <=> v4_waybel11(sK16)), introduced(avatar_definition, [new_symbols(naming, [spl60_10])])).
fof(f372047, plain, (! [X0] : (~ v4_waybel11(sK16) | v1_t_0topsp(sK15, X0, sK16) | ~ v4_waybel11(X0) | ~ m1_yellow_9(X0, sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(sK16)) | ~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(sK16))) | (~ spl60_9 | ~ spl60_12)), inference(resolution, [], [f1333, f1320])).
fof(f1320, plain, (m1_yellow_9(sK16, k2_yellow_2(sK14, sK14, sK15)) | ~ spl60_9), inference(avatar_component_clause, [], [f1318])).
fof(f1318, plain, (spl60_9 <=> m1_yellow_9(sK16, k2_yellow_2(sK14, sK14, sK15))), introduced(avatar_definition, [new_symbols(naming, [spl60_9])])).
fof(f1333, plain, (! [X6, X5] : (~ m1_yellow_9(X5, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X5) | v1_t_0topsp(sK15, X6, X5) | ~ v4_waybel11(X6) | ~ m1_yellow_9(X6, sK14) | ~ v1_funct_2(sK15, u1_struct_0(X6), u1_struct_0(X5)) | ~ m2_relset_1(sK15, u1_struct_0(X6), u1_struct_0(X5))) | ~ spl60_12), inference(avatar_component_clause, [], [f1332])).
fof(f1332, plain, (spl60_12 <=> ! [X5, X6] : (v1_t_0topsp(sK15, X6, X5) | ~ v4_waybel11(X5) | ~ m1_yellow_9(X5, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X6) | ~ m1_yellow_9(X6, sK14) | ~ v1_funct_2(sK15, u1_struct_0(X6), u1_struct_0(X5)) | ~ m2_relset_1(sK15, u1_struct_0(X6), u1_struct_0(X5)))), introduced(avatar_definition, [new_symbols(naming, [spl60_12])])).
fof(f371967, plain, (spl60_31635 | ~ spl60_9 | ~ spl60_66), inference(avatar_split_clause, [], [f371966, f3461, f1318, f371958])).
fof(f371958, plain, (spl60_31635 <=> l1_waybel_9(sK16)), introduced(avatar_definition, [new_symbols(naming, [spl60_31635])])).
fof(f3461, plain, (spl60_66 <=> l1_orders_2(k2_yellow_2(sK14, sK14, sK15))), introduced(avatar_definition, [new_symbols(naming, [spl60_66])])).
fof(f371966, plain, (l1_waybel_9(sK16) | (~ spl60_9 | ~ spl60_66)), inference(subsumption_resolution, [], [f371920, f3462])).
fof(f3462, plain, (l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ spl60_66), inference(avatar_component_clause, [], [f3461])).
fof(f371920, plain, (l1_waybel_9(sK16) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ spl60_9), inference(resolution, [], [f1320, f876])).
fof(f876, plain, ! [X0, X1] : (~ m1_yellow_9(X1, X0) | l1_waybel_9(X1) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f374])).
fof(f374, plain, ! [X0] : (! [X1] : (l1_waybel_9(X1) | ~ m1_yellow_9(X1, X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f91])).
fof(f91, plain, ! [X0] : (l1_orders_2(X0) => ! [X1] : (m1_yellow_9(X1, X0) => l1_waybel_9(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', dt_m1_yellow_9)).
fof(f371965, plain, (~ spl60_31635 | spl60_31636 | ~ spl60_9 | ~ spl60_66 | ~ spl60_67), inference(avatar_split_clause, [], [f371956, f3465, f3461, f1318, f371962, f371958])).
fof(f371956, plain, ((k2_yellow_2(sK14, sK14, sK15) = g1_orders_2(u1_struct_0(sK16), u1_orders_2(sK16))) | ~ l1_waybel_9(sK16) | (~ spl60_9 | ~ spl60_66 | ~ spl60_67)), inference(forward_demodulation, [], [f371955, f3467])).
fof(f371955, plain, ((g1_orders_2(u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_orders_2(k2_yellow_2(sK14, sK14, sK15))) = g1_orders_2(u1_struct_0(sK16), u1_orders_2(sK16))) | ~ l1_waybel_9(sK16) | (~ spl60_9 | ~ spl60_66)), inference(subsumption_resolution, [], [f371919, f3462])).
fof(f371919, plain, ((g1_orders_2(u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_orders_2(k2_yellow_2(sK14, sK14, sK15))) = g1_orders_2(u1_struct_0(sK16), u1_orders_2(sK16))) | ~ l1_waybel_9(sK16) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ spl60_9), inference(resolution, [], [f1320, f848])).
fof(f848, plain, ! [X0, X1] : (~ m1_yellow_9(X1, X0) | (g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1), u1_orders_2(X1))) | ~ l1_waybel_9(X1) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f530])).
fof(f530, plain, ! [X0] : (! [X1] : (((m1_yellow_9(X1, X0) | ~ (g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1), u1_orders_2(X1)))) & ((g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1), u1_orders_2(X1))) | ~ m1_yellow_9(X1, X0))) | ~ l1_waybel_9(X1)) | ~ l1_orders_2(X0)), inference(nnf_transformation, [], [f356])).
fof(f356, plain, ! [X0] : (! [X1] : ((m1_yellow_9(X1, X0) <=> (g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1), u1_orders_2(X1)))) | ~ l1_waybel_9(X1)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f72])).
fof(f72, plain, ! [X0] : (l1_orders_2(X0) => ! [X1] : (l1_waybel_9(X1) => (m1_yellow_9(X1, X0) <=> (g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1), u1_orders_2(X1)))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', d4_yellow_9)).
fof(f371907, plain, (spl60_31633 | ~ spl60_7), inference(avatar_split_clause, [], [f371906, f1308, f371898])).
fof(f371898, plain, (spl60_31633 <=> l1_waybel_9(sK17)), introduced(avatar_definition, [new_symbols(naming, [spl60_31633])])).
fof(f371906, plain, (l1_waybel_9(sK17) | ~ spl60_7), inference(subsumption_resolution, [], [f371861, f634])).
fof(f371861, plain, (l1_waybel_9(sK17) | ~ l1_orders_2(sK14) | ~ spl60_7), inference(resolution, [], [f1310, f876])).
fof(f371905, plain, (~ spl60_31633 | spl60_31634 | ~ spl60_7), inference(avatar_split_clause, [], [f371896, f1308, f371902, f371898])).
fof(f371896, plain, ((g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14)) = g1_orders_2(u1_struct_0(sK17), u1_orders_2(sK17))) | ~ l1_waybel_9(sK17) | ~ spl60_7), inference(subsumption_resolution, [], [f371860, f634])).
fof(f371860, plain, ((g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14)) = g1_orders_2(u1_struct_0(sK17), u1_orders_2(sK17))) | ~ l1_waybel_9(sK17) | ~ l1_orders_2(sK14) | ~ spl60_7), inference(resolution, [], [f1310, f848])).
fof(f371848, plain, (~ spl60_204 | spl60_12 | ~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200), inference(avatar_split_clause, [], [f371847, f8165, f8146, f8127, f4093, f4087, f4074, f3657, f3461, f1652, f1328, f1332, f8240])).
fof(f8240, plain, (spl60_204 <=> v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14)), introduced(avatar_definition, [new_symbols(naming, [spl60_204])])).
fof(f3657, plain, (spl60_70 <=> v2_funct_1(k3_waybel_1(sK14, sK14, sK15))), introduced(avatar_definition, [new_symbols(naming, [spl60_70])])).
fof(f4074, plain, (spl60_85 <=> v3_lattice3(k2_yellow_2(sK14, sK14, sK15))), introduced(avatar_definition, [new_symbols(naming, [spl60_85])])).
fof(f4087, plain, (spl60_87 <=> v1_lattice3(k2_yellow_2(sK14, sK14, sK15))), introduced(avatar_definition, [new_symbols(naming, [spl60_87])])).
fof(f4093, plain, (spl60_88 <=> v2_lattice3(k2_yellow_2(sK14, sK14, sK15))), introduced(avatar_definition, [new_symbols(naming, [spl60_88])])).
fof(f8127, plain, (spl60_198 <=> v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14)), introduced(avatar_definition, [new_symbols(naming, [spl60_198])])).
fof(f8146, plain, (spl60_199 <=> m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14))), introduced(avatar_definition, [new_symbols(naming, [spl60_199])])).
fof(f8165, plain, (spl60_200 <=> v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14))), introduced(avatar_definition, [new_symbols(naming, [spl60_200])])).
fof(f371847, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371846, f3856])).
fof(f3856, plain, (v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3855, f628])).
fof(f628, plain, v2_orders_2(sK14), inference(cnf_transformation, [], [f524])).
fof(f3855, plain, (v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3854, f634])).
fof(f3854, plain, (v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3847, f3830])).
fof(f3830, plain, (m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3829, f1653])).
fof(f3829, plain, (m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3828, f634])).
fof(f3828, plain, (m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3827, f1329])).
fof(f3827, plain, (m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f3826, f636])).
fof(f3826, plain, (m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(duplicate_literal_removal, [], [f3809])).
fof(f3809, plain, (m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(resolution, [], [f863, f1864])).
fof(f863, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f365])).
fof(f365, plain, ! [X0, X1, X2] : ((m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f364])).
fof(f364, plain, ! [X0, X1, X2] : ((m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f79])).
fof(f79, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', dt_k2_yellow_2)).
fof(f3847, plain, (v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ l1_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(resolution, [], [f3794, f814])).
fof(f814, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v2_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f337])).
fof(f337, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v2_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f336])).
fof(f336, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v2_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f62])).
fof(f62, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v2_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', cc6_yellow_0)).
fof(f3794, plain, (v4_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3793, f1653])).
fof(f3793, plain, (v4_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3792, f634])).
fof(f3792, plain, (v4_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3791, f1329])).
fof(f3791, plain, (v4_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f3790, f636])).
fof(f3790, plain, (v4_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(duplicate_literal_removal, [], [f3773])).
fof(f3773, plain, (v4_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(resolution, [], [f862, f1864])).
fof(f862, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f365])).
fof(f371846, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371845, f3853])).
fof(f3853, plain, (v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3852, f629])).
fof(f629, plain, v3_orders_2(sK14), inference(cnf_transformation, [], [f524])).
fof(f3852, plain, (v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3851, f634])).
fof(f3851, plain, (v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3846, f3830])).
fof(f3846, plain, (v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ l1_orders_2(sK14) | ~ v3_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(resolution, [], [f3794, f833])).
fof(f833, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v3_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v3_orders_2(X0)), inference(cnf_transformation, [], [f343])).
fof(f343, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v3_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v3_orders_2(X0)), inference(flattening, [], [f342])).
fof(f342, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v3_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v3_orders_2(X0))), inference(ennf_transformation, [], [f65])).
fof(f65, plain, ! [X0] : ((l1_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v3_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', cc7_yellow_0)).
fof(f371845, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371844, f3850])).
fof(f3850, plain, (v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3849, f630])).
fof(f630, plain, v4_orders_2(sK14), inference(cnf_transformation, [], [f524])).
fof(f3849, plain, (v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3848, f634])).
fof(f3848, plain, (v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v4_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3845, f3830])).
fof(f3845, plain, (v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ l1_orders_2(sK14) | ~ v4_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(resolution, [], [f3794, f839])).
fof(f839, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v4_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0)), inference(cnf_transformation, [], [f347])).
fof(f347, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v4_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0)), inference(flattening, [], [f346])).
fof(f346, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v4_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v4_orders_2(X0))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ! [X0] : ((l1_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v4_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', cc8_yellow_0)).
fof(f371844, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371843, f4089])).
fof(f4089, plain, (v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ spl60_87), inference(avatar_component_clause, [], [f4087])).
fof(f371843, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_88 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371842, f4095])).
fof(f4095, plain, (v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ spl60_88), inference(avatar_component_clause, [], [f4093])).
fof(f371842, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371841, f4075])).
fof(f4075, plain, (v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ spl60_85), inference(avatar_component_clause, [], [f4074])).
fof(f371841, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371840, f3462])).
fof(f371840, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371839, f628])).
fof(f371839, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371838, f629])).
fof(f371838, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371837, f630])).
fof(f371837, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371836, f631])).
fof(f631, plain, v1_lattice3(sK14), inference(cnf_transformation, [], [f524])).
fof(f371836, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371835, f632])).
fof(f632, plain, v2_lattice3(sK14), inference(cnf_transformation, [], [f524])).
fof(f371835, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371834, f633])).
fof(f633, plain, v3_lattice3(sK14), inference(cnf_transformation, [], [f524])).
fof(f371834, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371833, f634])).
fof(f371833, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371832, f3490])).
fof(f3490, plain, (v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3489, f1653])).
fof(f3489, plain, (v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3488, f634])).
fof(f3488, plain, (v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3487, f1329])).
fof(f3487, plain, (v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f3486, f636])).
fof(f3486, plain, (v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(duplicate_literal_removal, [], [f3469])).
fof(f3469, plain, (v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(resolution, [], [f864, f1864])).
fof(f864, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v1_funct_1(k3_waybel_1(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f367])).
fof(f367, plain, ! [X0, X1, X2] : ((m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f366])).
fof(f366, plain, ! [X0, X1, X2] : ((m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f81])).
fof(f81, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', dt_k3_waybel_1)).
fof(f371832, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199 | ~ spl60_200)), inference(subsumption_resolution, [], [f371831, f8167])).
fof(f8167, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ spl60_200), inference(avatar_component_clause, [], [f8165])).
fof(f371831, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_198 | ~ spl60_199)), inference(subsumption_resolution, [], [f371830, f8129])).
fof(f8129, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ spl60_198), inference(avatar_component_clause, [], [f8127])).
fof(f371830, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70 | ~ spl60_199)), inference(subsumption_resolution, [], [f371829, f8148])).
fof(f8148, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ spl60_199), inference(avatar_component_clause, [], [f8146])).
fof(f371829, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70)), inference(subsumption_resolution, [], [f371828, f3659])).
fof(f3659, plain, (v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ spl60_70), inference(avatar_component_clause, [], [f3657])).
fof(f371828, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8176, f1329])).
fof(f8176, plain, (! [X2, X3] : (~ m2_relset_1(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | v1_t_0topsp(sK15, X2, X3) | ~ v1_funct_2(sK15, u1_struct_0(X2), u1_struct_0(X3)) | ~ v1_funct_1(sK15) | ~ m1_yellow_9(X3, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X3) | ~ m1_yellow_9(X2, sK14) | ~ v4_waybel11(X2) | ~ v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26)), inference(superposition, [], [f1277, f5203])).
fof(f5203, plain, ((sK15 = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26)), inference(forward_demodulation, [], [f5202, f3872])).
fof(f5202, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ spl60_11), inference(subsumption_resolution, [], [f5201, f628])).
fof(f5201, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5200, f629])).
fof(f5200, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5199, f630])).
fof(f5199, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5198, f631])).
fof(f5198, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5197, f632])).
fof(f5197, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5196, f633])).
fof(f5196, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5195, f634])).
fof(f5195, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5194, f1329])).
fof(f5194, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(subsumption_resolution, [], [f5193, f636])).
fof(f5193, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(subsumption_resolution, [], [f5186, f637])).
fof(f637, plain, v7_waybel_1(sK15, sK14), inference(cnf_transformation, [], [f524])).
fof(f5186, plain, ((k2_waybel_1(sK14, sK14, sK15) = k1_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k3_waybel_1(sK14, sK14, sK15))) | ~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(resolution, [], [f1266, f638])).
fof(f1266, plain, ! [X0, X1] : (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | (k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f478])).
fof(f478, plain, ! [X0] : (! [X1] : (((k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) & v17_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f477])).
fof(f477, plain, ! [X0] : (! [X1] : (((k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) & v17_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) | (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f170])).
fof(f170, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => ((k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) & v17_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', t39_waybel34)).
fof(f1277, plain, ! [X4, X2, X0, X3, X1] : (~ m2_relset_1(k1_waybel34(X0, X1, X2), u1_struct_0(X3), u1_struct_0(X4)) | ~ v22_waybel_0(X2, X0, X1) | v1_t_0topsp(k1_waybel34(X0, X1, X2), X3, X4) | ~ v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_1(k1_waybel34(X0, X1, X2)) | ~ m1_yellow_9(X4, X0) | ~ v4_waybel11(X4) | ~ m1_yellow_9(X3, X1) | ~ v4_waybel11(X3) | ~ v2_funct_1(X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(equality_resolution, [], [f1260])).
fof(f1260, plain, ! [X4, X2, X0, X5, X3, X1] : (v1_t_0topsp(X5, X3, X4) | ~ v22_waybel_0(X2, X0, X1) | ~ (k1_waybel34(X0, X1, X2) = X5) | ~ m2_relset_1(X5, u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_2(X5, u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_1(X5) | ~ m1_yellow_9(X4, X0) | ~ v4_waybel11(X4) | ~ m1_yellow_9(X3, X1) | ~ v4_waybel11(X3) | ~ v2_funct_1(X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f626])).
fof(f626, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (((v22_waybel_0(X2, X0, X1) | ~ v1_t_0topsp(X5, X3, X4)) & (v1_t_0topsp(X5, X3, X4) | ~ v22_waybel_0(X2, X0, X1))) | ~ (k1_waybel34(X0, X1, X2) = X5) | ~ m2_relset_1(X5, u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_2(X5, u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_1(X5)) | ~ m1_yellow_9(X4, X0) | ~ v4_waybel11(X4)) | ~ m1_yellow_9(X3, X1) | ~ v4_waybel11(X3)) | ~ v2_funct_1(X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(nnf_transformation, [], [f474])).
fof(f474, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : ((v22_waybel_0(X2, X0, X1) <=> v1_t_0topsp(X5, X3, X4)) | ~ (k1_waybel34(X0, X1, X2) = X5) | ~ m2_relset_1(X5, u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_2(X5, u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_1(X5)) | ~ m1_yellow_9(X4, X0) | ~ v4_waybel11(X4)) | ~ m1_yellow_9(X3, X1) | ~ v4_waybel11(X3)) | ~ v2_funct_1(X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f473])).
fof(f473, plain, ! [X0] : (! [X1] : (! [X2] : ((! [X3] : (! [X4] : (! [X5] : (((v22_waybel_0(X2, X0, X1) <=> v1_t_0topsp(X5, X3, X4)) | ~ (k1_waybel34(X0, X1, X2) = X5)) | (~ m2_relset_1(X5, u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_2(X5, u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_1(X5))) | (~ m1_yellow_9(X4, X0) | ~ v4_waybel11(X4))) | (~ m1_yellow_9(X3, X1) | ~ v4_waybel11(X3))) | ~ v2_funct_1(X2)) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f168])).
fof(f168, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => (v2_funct_1(X2) => ! [X3] : ((m1_yellow_9(X3, X1) & v4_waybel11(X3)) => ! [X4] : ((m1_yellow_9(X4, X0) & v4_waybel11(X4)) => ! [X5] : ((m2_relset_1(X5, u1_struct_0(X3), u1_struct_0(X4)) & v1_funct_2(X5, u1_struct_0(X3), u1_struct_0(X4)) & v1_funct_1(X5)) => ((k1_waybel34(X0, X1, X2) = X5) => (v22_waybel_0(X2, X0, X1) <=> v1_t_0topsp(X5, X3, X4)))))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', t32_waybel34)).
fof(f371812, plain, (~ spl60_1 | ~ spl60_11 | spl60_204), inference(avatar_split_clause, [], [f371811, f8240, f1328, f1279])).
fof(f1279, plain, (spl60_1 <=> v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)), introduced(avatar_definition, [new_symbols(naming, [spl60_1])])).
fof(f371811, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | (~ spl60_11 | spl60_204)), inference(subsumption_resolution, [], [f371810, f628])).
fof(f371810, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_204)), inference(subsumption_resolution, [], [f371809, f629])).
fof(f371809, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_204)), inference(subsumption_resolution, [], [f371808, f630])).
fof(f371808, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_204)), inference(subsumption_resolution, [], [f371807, f631])).
fof(f371807, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_204)), inference(subsumption_resolution, [], [f371806, f632])).
fof(f371806, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_204)), inference(subsumption_resolution, [], [f371805, f633])).
fof(f371805, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_204)), inference(subsumption_resolution, [], [f371804, f634])).
fof(f371804, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_204)), inference(subsumption_resolution, [], [f371803, f1329])).
fof(f371803, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | spl60_204), inference(subsumption_resolution, [], [f371802, f636])).
fof(f371802, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | spl60_204), inference(subsumption_resolution, [], [f371801, f637])).
fof(f371801, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | spl60_204), inference(subsumption_resolution, [], [f24700, f638])).
fof(f24700, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | spl60_204), inference(resolution, [], [f8241, f1268])).
fof(f1268, plain, ! [X0, X1] : (v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) | ~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f627])).
fof(f627, plain, ! [X0] : (! [X1] : (((v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) | ~ v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)) & (v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) | ~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0))) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(nnf_transformation, [], [f481])).
fof(f481, plain, ! [X0] : (! [X1] : ((v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f480])).
fof(f480, plain, ! [X0] : (! [X1] : ((v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)) | (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f172])).
fof(f172, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', t40_waybel34)).
fof(f8241, plain, (~ v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | spl60_204), inference(avatar_component_clause, [], [f8240])).
fof(f371800, plain, (~ spl60_11 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_1829 | ~ spl60_2156 | ~ spl60_2175 | ~ spl60_2292), inference(avatar_contradiction_clause, [], [f371799])).
fof(f371799, plain, ($false | (~ spl60_11 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_1829 | ~ spl60_2156 | ~ spl60_2175 | ~ spl60_2292)), inference(subsumption_resolution, [], [f371798, f212708])).
fof(f212708, plain, (m1_yellow_9(sK58(sK14), sK14) | ~ spl60_1829), inference(subsumption_resolution, [], [f212670, f634])).
fof(f212670, plain, (~ l1_orders_2(sK14) | m1_yellow_9(sK58(sK14), sK14) | ~ spl60_1829), inference(equality_resolution, [], [f31188])).
fof(f31188, plain, (! [X1] : (~ (g1_orders_2(u1_struct_0(X1), u1_orders_2(X1)) = g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14))) | ~ l1_orders_2(X1) | m1_yellow_9(sK58(sK14), X1)) | ~ spl60_1829), inference(avatar_component_clause, [], [f31187])).
fof(f31187, plain, (spl60_1829 <=> ! [X1] : (~ (g1_orders_2(u1_struct_0(X1), u1_orders_2(X1)) = g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14))) | ~ l1_orders_2(X1) | m1_yellow_9(sK58(sK14), X1))), introduced(avatar_definition, [new_symbols(naming, [spl60_1829])])).
fof(f371798, plain, (~ m1_yellow_9(sK58(sK14), sK14) | (~ spl60_11 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_2156 | ~ spl60_2175 | ~ spl60_2292)), inference(subsumption_resolution, [], [f371797, f35711])).
fof(f35711, plain, (v4_waybel11(sK58(sK14)) | ~ spl60_2175), inference(avatar_component_clause, [], [f35710])).
fof(f35710, plain, (spl60_2175 <=> v4_waybel11(sK58(sK14))), introduced(avatar_definition, [new_symbols(naming, [spl60_2175])])).
fof(f371797, plain, (~ v4_waybel11(sK58(sK14)) | ~ m1_yellow_9(sK58(sK14), sK14) | (~ spl60_11 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_2156 | ~ spl60_2292)), inference(subsumption_resolution, [], [f371796, f4542])).
fof(f371796, plain, (~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v4_waybel11(sK58(sK14)) | ~ m1_yellow_9(sK58(sK14), sK14) | (~ spl60_11 | spl60_26 | ~ spl60_67 | ~ spl60_1818 | ~ spl60_2156 | ~ spl60_2292)), inference(subsumption_resolution, [], [f371792, f4568])).
fof(f371792, plain, (~ m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v4_waybel11(sK58(sK14)) | ~ m1_yellow_9(sK58(sK14), sK14) | (~ spl60_67 | ~ spl60_1818 | ~ spl60_2156 | ~ spl60_2292)), inference(superposition, [], [f266058, f171063])).
fof(f171063, plain, ((u1_struct_0(sK14) = u1_struct_0(sK58(sK14))) | ~ spl60_1818), inference(backward_demodulation, [], [f170798, f170797])).
fof(f170798, plain, ((u1_struct_0(sK26(sK14)) = u1_struct_0(sK58(sK14))) | ~ spl60_1818), inference(trivial_inequality_removal, [], [f170794])).
fof(f170794, plain, (~ (g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14)) = g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14))) | (u1_struct_0(sK26(sK14)) = u1_struct_0(sK58(sK14))) | ~ spl60_1818), inference(superposition, [], [f31094, f11908])).
fof(f11908, plain, (g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14)) = g1_orders_2(u1_struct_0(sK58(sK14)), u1_orders_2(sK58(sK14)))), inference(subsumption_resolution, [], [f11899, f634])).
fof(f11899, plain, (~ l1_orders_2(sK14) | (g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14)) = g1_orders_2(u1_struct_0(sK58(sK14)), u1_orders_2(sK58(sK14))))), inference(resolution, [], [f3021, f2794])).
fof(f2794, plain, sP13(sK14), inference(subsumption_resolution, [], [f2793, f628])).
fof(f2793, plain, (sP13(sK14) | ~ v2_orders_2(sK14)), inference(subsumption_resolution, [], [f2792, f629])).
fof(f2792, plain, (sP13(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(subsumption_resolution, [], [f2791, f630])).
fof(f2791, plain, (sP13(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(subsumption_resolution, [], [f2790, f631])).
fof(f2790, plain, (sP13(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(subsumption_resolution, [], [f2789, f632])).
fof(f2789, plain, (sP13(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(subsumption_resolution, [], [f2777, f633])).
fof(f2777, plain, (sP13(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(resolution, [], [f1251, f634])).
fof(f1251, plain, ! [X0] : (~ l1_orders_2(X0) | sP13(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f515])).
fof(f515, plain, ! [X0] : (sP13(X0) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(definition_folding, [], [f467, e514])).
fof(f514, plain, ! [X0] : (? [X1] : (v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_waybel11(X1) & v25_waybel_0(X1) & v24_waybel_0(X1) & v3_yellow_0(X1) & v2_yellow_0(X1) & v1_yellow_0(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & v2_pre_topc(X1) & ~ v3_struct_0(X1) & m1_yellow_9(X1, X0)) | ~ sP13(X0)), inference(usedef, [], [e514])).
fof(e514, plain, ! [X0] : (sP13(X0) <=> ? [X1] : (v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_waybel11(X1) & v25_waybel_0(X1) & v24_waybel_0(X1) & v3_yellow_0(X1) & v2_yellow_0(X1) & v1_yellow_0(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & v2_pre_topc(X1) & ~ v3_struct_0(X1) & m1_yellow_9(X1, X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f467, plain, ! [X0] : (? [X1] : (v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_waybel11(X1) & v25_waybel_0(X1) & v24_waybel_0(X1) & v3_yellow_0(X1) & v2_yellow_0(X1) & v1_yellow_0(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & v2_pre_topc(X1) & ~ v3_struct_0(X1) & m1_yellow_9(X1, X0)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f466])).
fof(f466, plain, ! [X0] : (? [X1] : (v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_waybel11(X1) & v25_waybel_0(X1) & v24_waybel_0(X1) & v3_yellow_0(X1) & v2_yellow_0(X1) & v1_yellow_0(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & v2_pre_topc(X1) & ~ v3_struct_0(X1) & m1_yellow_9(X1, X0)) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f180])).
fof(f180, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ? [X1] : (v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_waybel11(X1) & v25_waybel_0(X1) & v24_waybel_0(X1) & v3_yellow_0(X1) & v2_yellow_0(X1) & v1_yellow_0(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & v2_pre_topc(X1) & ~ v3_struct_0(X1) & m1_yellow_9(X1, X0))), inference(pure_predicate_removal, [], [f162])).
fof(f162, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ? [X1] : (v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_waybel11(X1) & v1_waybel_9(X1) & v25_waybel_0(X1) & v24_waybel_0(X1) & v3_yellow_0(X1) & v2_yellow_0(X1) & v1_yellow_0(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & v2_pre_topc(X1) & ~ v3_struct_0(X1) & m1_yellow_9(X1, X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', rc6_yellow_9)).
fof(f3021, plain, ! [X1] : (~ sP13(X1) | ~ l1_orders_2(X1) | (g1_orders_2(u1_struct_0(X1), u1_orders_2(X1)) = g1_orders_2(u1_struct_0(sK58(X1)), u1_orders_2(sK58(X1))))), inference(subsumption_resolution, [], [f3018, f1472])).
fof(f1472, plain, ! [X1] : (l1_waybel_9(sK58(X1)) | ~ l1_orders_2(X1) | ~ sP13(X1)), inference(resolution, [], [f876, f1236])).
fof(f1236, plain, ! [X0] : (m1_yellow_9(sK58(X0), X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f622])).
fof(f622, plain, ! [X0] : ((v3_lattice3(sK58(X0)) & v2_lattice3(sK58(X0)) & v1_lattice3(sK58(X0)) & v4_waybel11(sK58(X0)) & v25_waybel_0(sK58(X0)) & v24_waybel_0(sK58(X0)) & v3_yellow_0(sK58(X0)) & v2_yellow_0(sK58(X0)) & v1_yellow_0(sK58(X0)) & v4_orders_2(sK58(X0)) & v3_orders_2(sK58(X0)) & v2_orders_2(sK58(X0)) & v2_pre_topc(sK58(X0)) & ~ v3_struct_0(sK58(X0)) & m1_yellow_9(sK58(X0), X0)) | ~ sP13(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK58])], [f620, f621])).
fof(f621, plain, ! [X0] : (? [X1] : (v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_waybel11(X1) & v25_waybel_0(X1) & v24_waybel_0(X1) & v3_yellow_0(X1) & v2_yellow_0(X1) & v1_yellow_0(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & v2_pre_topc(X1) & ~ v3_struct_0(X1) & m1_yellow_9(X1, X0)) => (v3_lattice3(sK58(X0)) & v2_lattice3(sK58(X0)) & v1_lattice3(sK58(X0)) & v4_waybel11(sK58(X0)) & v25_waybel_0(sK58(X0)) & v24_waybel_0(sK58(X0)) & v3_yellow_0(sK58(X0)) & v2_yellow_0(sK58(X0)) & v1_yellow_0(sK58(X0)) & v4_orders_2(sK58(X0)) & v3_orders_2(sK58(X0)) & v2_orders_2(sK58(X0)) & v2_pre_topc(sK58(X0)) & ~ v3_struct_0(sK58(X0)) & m1_yellow_9(sK58(X0), X0))), introduced(choice_axiom, [])).
fof(f620, plain, ! [X0] : (? [X1] : (v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_waybel11(X1) & v25_waybel_0(X1) & v24_waybel_0(X1) & v3_yellow_0(X1) & v2_yellow_0(X1) & v1_yellow_0(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & v2_pre_topc(X1) & ~ v3_struct_0(X1) & m1_yellow_9(X1, X0)) | ~ sP13(X0)), inference(nnf_transformation, [], [f514])).
fof(f3018, plain, ! [X1] : ((g1_orders_2(u1_struct_0(X1), u1_orders_2(X1)) = g1_orders_2(u1_struct_0(sK58(X1)), u1_orders_2(sK58(X1)))) | ~ l1_waybel_9(sK58(X1)) | ~ l1_orders_2(X1) | ~ sP13(X1)), inference(resolution, [], [f848, f1236])).
fof(f266058, plain, (! [X1] : (~ m2_relset_1(sK15, u1_struct_0(X1), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(X1), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X1, sK14)) | (~ spl60_67 | ~ spl60_2156 | ~ spl60_2292)), inference(forward_demodulation, [], [f266031, f266025])).
fof(f266031, plain, (! [X1] : (~ v1_funct_2(sK15, u1_struct_0(X1), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v4_waybel11(X1) | ~ m2_relset_1(sK15, u1_struct_0(X1), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | ~ m1_yellow_9(X1, sK14)) | (~ spl60_67 | ~ spl60_2156 | ~ spl60_2292)), inference(backward_demodulation, [], [f39741, f266025])).
fof(f39741, plain, (! [X1] : (~ v4_waybel11(X1) | ~ v1_funct_2(sK15, u1_struct_0(X1), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | ~ m2_relset_1(sK15, u1_struct_0(X1), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | ~ m1_yellow_9(X1, sK14)) | ~ spl60_2292), inference(avatar_component_clause, [], [f39740])).
fof(f39740, plain, (spl60_2292 <=> ! [X1] : (~ v4_waybel11(X1) | ~ v1_funct_2(sK15, u1_struct_0(X1), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | ~ m2_relset_1(sK15, u1_struct_0(X1), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | ~ m1_yellow_9(X1, sK14))), introduced(avatar_definition, [new_symbols(naming, [spl60_2292])])).
fof(f265972, plain, (~ spl60_2145 | spl60_2155), inference(avatar_contradiction_clause, [], [f265971])).
fof(f265971, plain, ($false | (~ spl60_2145 | spl60_2155)), inference(subsumption_resolution, [], [f265970, f35127])).
fof(f35127, plain, (l1_orders_2(sK58(k2_yellow_2(sK14, sK14, sK15))) | ~ spl60_2145), inference(avatar_component_clause, [], [f35126])).
fof(f35126, plain, (spl60_2145 <=> l1_orders_2(sK58(k2_yellow_2(sK14, sK14, sK15)))), introduced(avatar_definition, [new_symbols(naming, [spl60_2145])])).
fof(f265970, plain, (~ l1_orders_2(sK58(k2_yellow_2(sK14, sK14, sK15))) | spl60_2155), inference(resolution, [], [f35171, f1886])).
fof(f1886, plain, ! [X0] : (m1_relset_1(u1_orders_2(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ l1_orders_2(X0)), inference(resolution, [], [f878, f1255])).
fof(f878, plain, ! [X0] : (m2_relset_1(u1_orders_2(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f376])).
fof(f376, plain, ! [X0] : (m2_relset_1(u1_orders_2(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f93])).
fof(f93, plain, ! [X0] : (l1_orders_2(X0) => m2_relset_1(u1_orders_2(X0), u1_struct_0(X0), u1_struct_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', dt_u1_orders_2)).
fof(f35171, plain, (~ m1_relset_1(u1_orders_2(sK58(k2_yellow_2(sK14, sK14, sK15))), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | spl60_2155), inference(avatar_component_clause, [], [f35169])).
fof(f35169, plain, (spl60_2155 <=> m1_relset_1(u1_orders_2(sK58(k2_yellow_2(sK14, sK14, sK15))), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))))), introduced(avatar_definition, [new_symbols(naming, [spl60_2155])])).
fof(f119268, plain, (~ spl60_1788 | spl60_1817), inference(avatar_contradiction_clause, [], [f119267])).
fof(f119267, plain, ($false | (~ spl60_1788 | spl60_1817)), inference(subsumption_resolution, [], [f119266, f30945])).
fof(f30945, plain, (l1_orders_2(sK26(sK14)) | ~ spl60_1788), inference(avatar_component_clause, [], [f30944])).
fof(f30944, plain, (spl60_1788 <=> l1_orders_2(sK26(sK14))), introduced(avatar_definition, [new_symbols(naming, [spl60_1788])])).
fof(f119266, plain, (~ l1_orders_2(sK26(sK14)) | spl60_1817), inference(resolution, [], [f31091, f1886])).
fof(f31091, plain, (~ m1_relset_1(u1_orders_2(sK26(sK14)), u1_struct_0(sK26(sK14)), u1_struct_0(sK26(sK14))) | spl60_1817), inference(avatar_component_clause, [], [f31089])).
fof(f31089, plain, (spl60_1817 <=> m1_relset_1(u1_orders_2(sK26(sK14)), u1_struct_0(sK26(sK14)), u1_struct_0(sK26(sK14)))), introduced(avatar_definition, [new_symbols(naming, [spl60_1817])])).
fof(f51778, plain, (~ spl60_73 | spl60_2291), inference(avatar_contradiction_clause, [], [f51777])).
fof(f51777, plain, ($false | (~ spl60_73 | spl60_2291)), inference(subsumption_resolution, [], [f51776, f3882])).
fof(f3882, plain, (sP13(k2_yellow_2(sK14, sK14, sK15)) | ~ spl60_73), inference(avatar_component_clause, [], [f3881])).
fof(f3881, plain, (spl60_73 <=> sP13(k2_yellow_2(sK14, sK14, sK15))), introduced(avatar_definition, [new_symbols(naming, [spl60_73])])).
fof(f51776, plain, (~ sP13(k2_yellow_2(sK14, sK14, sK15)) | spl60_2291), inference(resolution, [], [f39738, f1247])).
fof(f1247, plain, ! [X0] : (v4_waybel11(sK58(X0)) | ~ sP13(X0)), inference(cnf_transformation, [], [f622])).
fof(f39738, plain, (~ v4_waybel11(sK58(k2_yellow_2(sK14, sK14, sK15))) | spl60_2291), inference(avatar_component_clause, [], [f39736])).
fof(f39736, plain, (spl60_2291 <=> v4_waybel11(sK58(k2_yellow_2(sK14, sK14, sK15)))), introduced(avatar_definition, [new_symbols(naming, [spl60_2291])])).
fof(f51196, plain, (~ spl60_66 | ~ spl60_73 | spl60_2145), inference(avatar_contradiction_clause, [], [f51195])).
fof(f51195, plain, ($false | (~ spl60_66 | ~ spl60_73 | spl60_2145)), inference(subsumption_resolution, [], [f51194, f3462])).
fof(f51194, plain, (~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_73 | spl60_2145)), inference(subsumption_resolution, [], [f51193, f3882])).
fof(f51193, plain, (~ sP13(k2_yellow_2(sK14, sK14, sK15)) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | spl60_2145), inference(resolution, [], [f35128, f1693])).
fof(f1693, plain, ! [X0] : (l1_orders_2(sK58(X0)) | ~ sP13(X0) | ~ l1_orders_2(X0)), inference(resolution, [], [f1472, f874])).
fof(f874, plain, ! [X0] : (~ l1_waybel_9(X0) | l1_orders_2(X0)), inference(cnf_transformation, [], [f372])).
fof(f372, plain, ! [X0] : ((l1_orders_2(X0) & l1_pre_topc(X0)) | ~ l1_waybel_9(X0)), inference(ennf_transformation, [], [f87])).
fof(f87, plain, ! [X0] : (l1_waybel_9(X0) => (l1_orders_2(X0) & l1_pre_topc(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', dt_l1_waybel_9)).
fof(f35128, plain, (~ l1_orders_2(sK58(k2_yellow_2(sK14, sK14, sK15))) | spl60_2145), inference(avatar_component_clause, [], [f35126])).
fof(f39742, plain, (~ spl60_2291 | spl60_2292 | ~ spl60_73 | ~ spl60_205), inference(avatar_split_clause, [], [f39734, f8244, f3881, f39740, f39736])).
fof(f8244, plain, (spl60_205 <=> ! [X1, X0] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v4_waybel11(X0) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)))), introduced(avatar_definition, [new_symbols(naming, [spl60_205])])).
fof(f39734, plain, (! [X1] : (~ v4_waybel11(X1) | ~ m1_yellow_9(X1, sK14) | ~ v4_waybel11(sK58(k2_yellow_2(sK14, sK14, sK15))) | ~ m2_relset_1(sK15, u1_struct_0(X1), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | ~ v1_funct_2(sK15, u1_struct_0(X1), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))))) | (~ spl60_73 | ~ spl60_205)), inference(subsumption_resolution, [], [f39732, f3882])).
fof(f39732, plain, (! [X1] : (~ v4_waybel11(X1) | ~ m1_yellow_9(X1, sK14) | ~ v4_waybel11(sK58(k2_yellow_2(sK14, sK14, sK15))) | ~ m2_relset_1(sK15, u1_struct_0(X1), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | ~ v1_funct_2(sK15, u1_struct_0(X1), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15)))) | ~ sP13(k2_yellow_2(sK14, sK14, sK15))) | ~ spl60_205), inference(resolution, [], [f8245, f1236])).
fof(f8245, plain, (! [X0, X1] : (~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X0) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X1) | ~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1))) | ~ spl60_205), inference(avatar_component_clause, [], [f8244])).
fof(f36138, plain, spl60_2175, inference(avatar_contradiction_clause, [], [f36137])).
fof(f36137, plain, ($false | spl60_2175), inference(subsumption_resolution, [], [f36136, f2794])).
fof(f36136, plain, (~ sP13(sK14) | spl60_2175), inference(resolution, [], [f35712, f1247])).
fof(f35712, plain, (~ v4_waybel11(sK58(sK14)) | spl60_2175), inference(avatar_component_clause, [], [f35710])).
fof(f35175, plain, (~ spl60_2155 | spl60_2156 | ~ spl60_66 | ~ spl60_67 | ~ spl60_73), inference(avatar_split_clause, [], [f35109, f3881, f3465, f3461, f35173, f35169])).
fof(f35109, plain, (! [X2, X3] : (~ (g1_orders_2(X2, X3) = k2_yellow_2(sK14, sK14, sK15)) | (u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))) = X2) | ~ m1_relset_1(u1_orders_2(sK58(k2_yellow_2(sK14, sK14, sK15))), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))), u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))))) | (~ spl60_66 | ~ spl60_67 | ~ spl60_73)), inference(superposition, [], [f1052, f11907])).
fof(f11907, plain, ((k2_yellow_2(sK14, sK14, sK15) = g1_orders_2(u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))), u1_orders_2(sK58(k2_yellow_2(sK14, sK14, sK15))))) | (~ spl60_66 | ~ spl60_67 | ~ spl60_73)), inference(forward_demodulation, [], [f11906, f3467])).
fof(f11906, plain, ((g1_orders_2(u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_orders_2(k2_yellow_2(sK14, sK14, sK15))) = g1_orders_2(u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))), u1_orders_2(sK58(k2_yellow_2(sK14, sK14, sK15))))) | (~ spl60_66 | ~ spl60_73)), inference(subsumption_resolution, [], [f11898, f3462])).
fof(f11898, plain, (~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (g1_orders_2(u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_orders_2(k2_yellow_2(sK14, sK14, sK15))) = g1_orders_2(u1_struct_0(sK58(k2_yellow_2(sK14, sK14, sK15))), u1_orders_2(sK58(k2_yellow_2(sK14, sK14, sK15))))) | ~ spl60_73), inference(resolution, [], [f3021, f3882])).
fof(f1052, plain, ! [X2, X0, X3, X1] : (~ (g1_orders_2(X0, X1) = g1_orders_2(X2, X3)) | (X0 = X2) | ~ m1_relset_1(X1, X0, X0)), inference(cnf_transformation, [], [f429])).
fof(f429, plain, ! [X0, X1] : (! [X2, X3] : (((X1 = X3) & (X0 = X2)) | ~ (g1_orders_2(X0, X1) = g1_orders_2(X2, X3))) | ~ m1_relset_1(X1, X0, X0)), inference(ennf_transformation, [], [f131])).
fof(f131, plain, ! [X0, X1] : (m1_relset_1(X1, X0, X0) => ! [X2, X3] : ((g1_orders_2(X0, X1) = g1_orders_2(X2, X3)) => ((X1 = X3) & (X0 = X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', free_g1_orders_2)).
fof(f32302, plain, spl60_1828, inference(avatar_contradiction_clause, [], [f32301])).
fof(f32301, plain, ($false | spl60_1828), inference(subsumption_resolution, [], [f32300, f2794])).
fof(f32300, plain, (~ sP13(sK14) | spl60_1828), inference(subsumption_resolution, [], [f32299, f634])).
fof(f32299, plain, (~ l1_orders_2(sK14) | ~ sP13(sK14) | spl60_1828), inference(resolution, [], [f31185, f1472])).
fof(f31185, plain, (~ l1_waybel_9(sK58(sK14)) | spl60_1828), inference(avatar_component_clause, [], [f31183])).
fof(f31183, plain, (spl60_1828 <=> l1_waybel_9(sK58(sK14))), introduced(avatar_definition, [new_symbols(naming, [spl60_1828])])).
fof(f31189, plain, (~ spl60_1828 | spl60_1829), inference(avatar_split_clause, [], [f31134, f31187, f31183])).
fof(f31134, plain, ! [X1] : (~ (g1_orders_2(u1_struct_0(X1), u1_orders_2(X1)) = g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14))) | m1_yellow_9(sK58(sK14), X1) | ~ l1_waybel_9(sK58(sK14)) | ~ l1_orders_2(X1)), inference(superposition, [], [f849, f11908])).
fof(f849, plain, ! [X0, X1] : (~ (g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = g1_orders_2(u1_struct_0(X1), u1_orders_2(X1))) | m1_yellow_9(X1, X0) | ~ l1_waybel_9(X1) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f530])).
fof(f31132, plain, spl60_1788, inference(avatar_contradiction_clause, [], [f31131])).
fof(f31131, plain, ($false | spl60_1788), inference(subsumption_resolution, [], [f31130, f634])).
fof(f31130, plain, (~ l1_orders_2(sK14) | spl60_1788), inference(resolution, [], [f30946, f1476])).
fof(f1476, plain, ! [X0] : (l1_orders_2(sK26(X0)) | ~ l1_orders_2(X0)), inference(resolution, [], [f1473, f874])).
fof(f1473, plain, ! [X0] : (l1_waybel_9(sK26(X0)) | ~ l1_orders_2(X0)), inference(duplicate_literal_removal, [], [f1471])).
fof(f1471, plain, ! [X0] : (l1_waybel_9(sK26(X0)) | ~ l1_orders_2(X0) | ~ l1_orders_2(X0)), inference(resolution, [], [f876, f886])).
fof(f886, plain, ! [X0] : (m1_yellow_9(sK26(X0), X0) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f546])).
fof(f546, plain, ! [X0] : (m1_yellow_9(sK26(X0), X0) | ~ l1_orders_2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK26])], [f378, f545])).
fof(f545, plain, ! [X0] : (? [X1] : m1_yellow_9(X1, X0) => m1_yellow_9(sK26(X0), X0)), introduced(choice_axiom, [])).
fof(f378, plain, ! [X0] : (? [X1] : m1_yellow_9(X1, X0) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f102])).
fof(f102, plain, ! [X0] : (l1_orders_2(X0) => ? [X1] : m1_yellow_9(X1, X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', existence_m1_yellow_9)).
fof(f30946, plain, (~ l1_orders_2(sK26(sK14)) | spl60_1788), inference(avatar_component_clause, [], [f30944])).
fof(f31095, plain, (~ spl60_1817 | spl60_1818), inference(avatar_split_clause, [], [f30933, f31093, f31089])).
fof(f30933, plain, ! [X2, X3] : (~ (g1_orders_2(X2, X3) = g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14))) | (u1_struct_0(sK26(sK14)) = X2) | ~ m1_relset_1(u1_orders_2(sK26(sK14)), u1_struct_0(sK26(sK14)), u1_struct_0(sK26(sK14)))), inference(superposition, [], [f1052, f10822])).
fof(f10822, plain, (g1_orders_2(u1_struct_0(sK14), u1_orders_2(sK14)) = g1_orders_2(u1_struct_0(sK26(sK14)), u1_orders_2(sK26(sK14)))), inference(resolution, [], [f3020, f634])).
fof(f3020, plain, ! [X0] : (~ l1_orders_2(X0) | (g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = g1_orders_2(u1_struct_0(sK26(X0)), u1_orders_2(sK26(X0))))), inference(subsumption_resolution, [], [f3019, f1473])).
fof(f3019, plain, ! [X0] : ((g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = g1_orders_2(u1_struct_0(sK26(X0)), u1_orders_2(sK26(X0)))) | ~ l1_waybel_9(sK26(X0)) | ~ l1_orders_2(X0)), inference(duplicate_literal_removal, [], [f3017])).
fof(f3017, plain, ! [X0] : ((g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = g1_orders_2(u1_struct_0(sK26(X0)), u1_orders_2(sK26(X0)))) | ~ l1_waybel_9(sK26(X0)) | ~ l1_orders_2(X0) | ~ l1_orders_2(X0)), inference(resolution, [], [f848, f886])).
fof(f24575, plain, (spl60_1 | ~ spl60_11 | ~ spl60_204), inference(avatar_contradiction_clause, [], [f24574])).
fof(f24574, plain, ($false | (spl60_1 | ~ spl60_11 | ~ spl60_204)), inference(subsumption_resolution, [], [f24573, f628])).
fof(f24573, plain, (~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_11 | ~ spl60_204)), inference(subsumption_resolution, [], [f24572, f629])).
fof(f24572, plain, (~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_11 | ~ spl60_204)), inference(subsumption_resolution, [], [f24571, f630])).
fof(f24571, plain, (~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_11 | ~ spl60_204)), inference(subsumption_resolution, [], [f24570, f631])).
fof(f24570, plain, (~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_11 | ~ spl60_204)), inference(subsumption_resolution, [], [f24569, f632])).
fof(f24569, plain, (~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_11 | ~ spl60_204)), inference(subsumption_resolution, [], [f24568, f633])).
fof(f24568, plain, (~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_11 | ~ spl60_204)), inference(subsumption_resolution, [], [f24567, f634])).
fof(f24567, plain, (~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_11 | ~ spl60_204)), inference(subsumption_resolution, [], [f24566, f1329])).
fof(f24566, plain, (~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_204)), inference(subsumption_resolution, [], [f24565, f636])).
fof(f24565, plain, (~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_204)), inference(subsumption_resolution, [], [f24564, f637])).
fof(f24564, plain, (~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_204)), inference(subsumption_resolution, [], [f24563, f638])).
fof(f24563, plain, (~ m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (spl60_1 | ~ spl60_204)), inference(subsumption_resolution, [], [f24562, f1281])).
fof(f1281, plain, (~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | spl60_1), inference(avatar_component_clause, [], [f1279])).
fof(f24562, plain, (v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_204), inference(resolution, [], [f8242, f1269])).
fof(f1269, plain, ! [X0, X1] : (~ v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) | v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f627])).
fof(f8242, plain, (v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ spl60_204), inference(avatar_component_clause, [], [f8240])).
fof(f17172, plain, (spl60_194 | ~ spl60_11 | spl60_26), inference(avatar_split_clause, [], [f17171, f1652, f1328, f8057])).
fof(f8057, plain, (spl60_194 <=> m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)))), introduced(avatar_definition, [new_symbols(naming, [spl60_194])])).
fof(f17171, plain, (m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f17170, f1864])).
fof(f17170, plain, (m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f17169, f1653])).
fof(f17169, plain, (m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | v3_struct_0(sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f17168, f634])).
fof(f17168, plain, (m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f17167, f1329])).
fof(f17167, plain, (m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f17017, f636])).
fof(f17017, plain, (m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | (~ spl60_11 | spl60_26)), inference(duplicate_literal_removal, [], [f17006])).
fof(f17006, plain, (m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | (~ spl60_11 | spl60_26)), inference(superposition, [], [f4549, f3872])).
fof(f4549, plain, ! [X14, X12, X13] : (m1_relset_1(k2_waybel_1(X13, X14, X12), u1_struct_0(X13), u1_struct_0(k2_yellow_2(X13, X14, X12))) | ~ v1_funct_2(X12, u1_struct_0(X13), u1_struct_0(X14)) | ~ v1_funct_1(X12) | ~ l1_orders_2(X14) | v3_struct_0(X14) | ~ l1_orders_2(X13) | v3_struct_0(X13) | ~ m1_relset_1(X12, u1_struct_0(X13), u1_struct_0(X14))), inference(resolution, [], [f860, f1255])).
fof(f8246, plain, (~ spl60_200 | ~ spl60_198 | ~ spl60_199 | spl60_204 | spl60_205 | ~ spl60_11 | ~ spl60_12 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88), inference(avatar_split_clause, [], [f8238, f4093, f4087, f4074, f3657, f3461, f1652, f1332, f1328, f8244, f8240, f8146, f8127, f8165])).
fof(f8238, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14))) | (~ spl60_11 | ~ spl60_12 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8237, f1333])).
fof(f8237, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8236, f3856])).
fof(f8236, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8235, f3853])).
fof(f8235, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8234, f3850])).
fof(f8234, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8233, f4089])).
fof(f8233, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85 | ~ spl60_88)), inference(subsumption_resolution, [], [f8232, f4095])).
fof(f8232, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70 | ~ spl60_85)), inference(subsumption_resolution, [], [f8231, f4075])).
fof(f8231, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_70)), inference(subsumption_resolution, [], [f8230, f3462])).
fof(f8230, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70)), inference(subsumption_resolution, [], [f8229, f628])).
fof(f8229, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70)), inference(subsumption_resolution, [], [f8228, f629])).
fof(f8228, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70)), inference(subsumption_resolution, [], [f8227, f630])).
fof(f8227, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70)), inference(subsumption_resolution, [], [f8226, f631])).
fof(f8226, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70)), inference(subsumption_resolution, [], [f8225, f632])).
fof(f8225, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70)), inference(subsumption_resolution, [], [f8224, f633])).
fof(f8224, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70)), inference(subsumption_resolution, [], [f8223, f634])).
fof(f8223, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70)), inference(subsumption_resolution, [], [f8222, f3490])).
fof(f8222, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_70)), inference(subsumption_resolution, [], [f8221, f3659])).
fof(f8221, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8175, f1329])).
fof(f8175, plain, (! [X0, X1] : (~ m2_relset_1(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_t_0topsp(sK15, X0, X1) | v22_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(sK15) | ~ m1_yellow_9(X1, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X1) | ~ m1_yellow_9(X0, sK14) | ~ v4_waybel11(X0) | ~ v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ v1_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26)), inference(superposition, [], [f1276, f5203])).
fof(f1276, plain, ! [X4, X2, X0, X3, X1] : (~ m2_relset_1(k1_waybel34(X0, X1, X2), u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_t_0topsp(k1_waybel34(X0, X1, X2), X3, X4) | v22_waybel_0(X2, X0, X1) | ~ v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_1(k1_waybel34(X0, X1, X2)) | ~ m1_yellow_9(X4, X0) | ~ v4_waybel11(X4) | ~ m1_yellow_9(X3, X1) | ~ v4_waybel11(X3) | ~ v2_funct_1(X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(equality_resolution, [], [f1261])).
fof(f1261, plain, ! [X4, X2, X0, X5, X3, X1] : (v22_waybel_0(X2, X0, X1) | ~ v1_t_0topsp(X5, X3, X4) | ~ (k1_waybel34(X0, X1, X2) = X5) | ~ m2_relset_1(X5, u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_2(X5, u1_struct_0(X3), u1_struct_0(X4)) | ~ v1_funct_1(X5) | ~ m1_yellow_9(X4, X0) | ~ v4_waybel11(X4) | ~ m1_yellow_9(X3, X1) | ~ v4_waybel11(X3) | ~ v2_funct_1(X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f626])).
fof(f8168, plain, (~ spl60_194 | spl60_200 | ~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_87 | ~ spl60_88), inference(avatar_split_clause, [], [f8163, f4093, f4087, f3461, f1652, f1328, f8165, f8057])).
fof(f8163, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8162, f3856])).
fof(f8162, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8161, f3853])).
fof(f8161, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8160, f3850])).
fof(f8160, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8159, f4089])).
fof(f8159, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_88)), inference(subsumption_resolution, [], [f8158, f4095])).
fof(f8158, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66)), inference(subsumption_resolution, [], [f8157, f3462])).
fof(f8157, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8156, f628])).
fof(f8156, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8155, f629])).
fof(f8155, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8154, f630])).
fof(f8154, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8153, f631])).
fof(f8153, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8152, f632])).
fof(f8152, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8151, f634])).
fof(f8151, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ l1_orders_2(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8150, f1329])).
fof(f8150, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8038, f4542])).
fof(f8038, plain, (v1_funct_2(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(superposition, [], [f856, f5183])).
fof(f5183, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(forward_demodulation, [], [f5182, f3872])).
fof(f5182, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ spl60_11), inference(subsumption_resolution, [], [f5181, f628])).
fof(f5181, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5180, f629])).
fof(f5180, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5179, f630])).
fof(f5179, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5178, f631])).
fof(f5178, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5177, f632])).
fof(f5177, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5176, f633])).
fof(f5176, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5175, f634])).
fof(f5175, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f5174, f1329])).
fof(f5174, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(subsumption_resolution, [], [f5173, f636])).
fof(f5173, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(subsumption_resolution, [], [f5166, f637])).
fof(f5166, plain, ((k3_waybel_1(sK14, sK14, sK15) = k2_waybel34(k2_yellow_2(sK14, sK14, sK15), sK14, k2_waybel_1(sK14, sK14, sK15))) | ~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(resolution, [], [f1265, f638])).
fof(f1265, plain, ! [X0, X1] : (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f478])).
fof(f856, plain, ! [X2, X0, X1] : (v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f361])).
fof(f361, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f360])).
fof(f360, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f77])).
fof(f77, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', dt_k2_waybel34)).
fof(f8149, plain, (~ spl60_194 | spl60_199 | ~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_87 | ~ spl60_88), inference(avatar_split_clause, [], [f8144, f4093, f4087, f3461, f1652, f1328, f8146, f8057])).
fof(f8144, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8143, f3856])).
fof(f8143, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8142, f3853])).
fof(f8142, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8141, f3850])).
fof(f8141, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8140, f4089])).
fof(f8140, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_88)), inference(subsumption_resolution, [], [f8139, f4095])).
fof(f8139, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66)), inference(subsumption_resolution, [], [f8138, f3462])).
fof(f8138, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8137, f628])).
fof(f8137, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8136, f629])).
fof(f8136, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8135, f630])).
fof(f8135, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8134, f631])).
fof(f8134, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8133, f632])).
fof(f8133, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8132, f634])).
fof(f8132, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ l1_orders_2(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8131, f1329])).
fof(f8131, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8037, f4542])).
fof(f8037, plain, (m2_relset_1(k3_waybel_1(sK14, sK14, sK15), u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_struct_0(sK14)) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(superposition, [], [f857, f5183])).
fof(f857, plain, ! [X2, X0, X1] : (m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f361])).
fof(f8130, plain, (~ spl60_194 | spl60_198 | ~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88), inference(avatar_split_clause, [], [f8125, f4093, f4087, f4074, f3461, f1652, f1328, f8127, f8057])).
fof(f8125, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8124, f3856])).
fof(f8124, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8123, f3853])).
fof(f8123, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8122, f3850])).
fof(f8122, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_85 | ~ spl60_87 | ~ spl60_88)), inference(subsumption_resolution, [], [f8121, f4089])).
fof(f8121, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_85 | ~ spl60_88)), inference(subsumption_resolution, [], [f8120, f4095])).
fof(f8120, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_85)), inference(subsumption_resolution, [], [f8119, f4075])).
fof(f8119, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66)), inference(subsumption_resolution, [], [f8118, f3462])).
fof(f8118, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8117, f628])).
fof(f8117, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8116, f629])).
fof(f8116, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8115, f630])).
fof(f8115, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8114, f631])).
fof(f8114, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8113, f632])).
fof(f8113, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8112, f633])).
fof(f8112, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8111, f634])).
fof(f8111, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8110, f1329])).
fof(f8110, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8109, f4542])).
fof(f8109, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f8036, f4870])).
fof(f4870, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4869, f628])).
fof(f4869, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4868, f629])).
fof(f4868, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4867, f630])).
fof(f4867, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4866, f631])).
fof(f4866, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4865, f632])).
fof(f4865, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4864, f633])).
fof(f4864, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4863, f634])).
fof(f4863, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4862, f1329])).
fof(f4862, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4861, f636])).
fof(f4861, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4860, f637])).
fof(f4860, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4859, f638])).
fof(f4859, plain, (v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ m2_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(superposition, [], [f1263, f3872])).
fof(f1263, plain, ! [X0, X1] : (v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f478])).
fof(f8036, plain, (v17_waybel_0(k3_waybel_1(sK14, sK14, sK15), k2_yellow_2(sK14, sK14, sK15), sK14) | ~ m1_relset_1(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v18_waybel_0(sK15, sK14, k2_yellow_2(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(k2_yellow_2(sK14, sK14, sK15))) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(superposition, [], [f945, f5183])).
fof(f945, plain, ! [X2, X0, X1] : (v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f396])).
fof(f396, plain, ! [X0, X1, X2] : ((v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f395])).
fof(f395, plain, ! [X0, X1, X2] : ((v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f211])).
fof(f211, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f195])).
fof(f195, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_partfun1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f112])).
fof(f112, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v21_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_partfun1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', fc2_waybel34)).
fof(f4631, plain, (~ spl60_87 | ~ spl60_85 | spl60_73 | ~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_88), inference(avatar_split_clause, [], [f4630, f4093, f3461, f1652, f1328, f3881, f4074, f4087])).
fof(f4630, plain, (sP13(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_88)), inference(subsumption_resolution, [], [f4629, f3856])).
fof(f4629, plain, (sP13(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_88)), inference(subsumption_resolution, [], [f4628, f3853])).
fof(f4628, plain, (sP13(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66 | ~ spl60_88)), inference(subsumption_resolution, [], [f4627, f3850])).
fof(f4627, plain, (sP13(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_66 | ~ spl60_88)), inference(subsumption_resolution, [], [f4066, f4095])).
fof(f4066, plain, (sP13(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ spl60_66), inference(resolution, [], [f3462, f1251])).
fof(f4615, plain, (~ spl60_11 | spl60_26 | spl60_85), inference(avatar_contradiction_clause, [], [f4614])).
fof(f4614, plain, ($false | (~ spl60_11 | spl60_26 | spl60_85)), inference(subsumption_resolution, [], [f4613, f628])).
fof(f4613, plain, (~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26 | spl60_85)), inference(subsumption_resolution, [], [f4612, f629])).
fof(f4612, plain, (~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26 | spl60_85)), inference(subsumption_resolution, [], [f4611, f630])).
fof(f4611, plain, (~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26 | spl60_85)), inference(subsumption_resolution, [], [f4610, f631])).
fof(f4610, plain, (~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26 | spl60_85)), inference(subsumption_resolution, [], [f4609, f632])).
fof(f4609, plain, (~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26 | spl60_85)), inference(subsumption_resolution, [], [f4608, f633])).
fof(f4608, plain, (~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26 | spl60_85)), inference(subsumption_resolution, [], [f4607, f634])).
fof(f4607, plain, (~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26 | spl60_85)), inference(subsumption_resolution, [], [f4606, f1329])).
fof(f4606, plain, (~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26 | spl60_85)), inference(subsumption_resolution, [], [f4605, f636])).
fof(f4605, plain, (~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | (~ spl60_11 | spl60_26 | spl60_85)), inference(subsumption_resolution, [], [f4604, f3322])).
fof(f3322, plain, (v6_waybel_1(sK15, sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3321, f1653])).
fof(f3321, plain, (v6_waybel_1(sK15, sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3320, f634])).
fof(f3320, plain, (v6_waybel_1(sK15, sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3319, f1329])).
fof(f3319, plain, (~ v1_funct_1(sK15) | v6_waybel_1(sK15, sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f3318, f636])).
fof(f3318, plain, (~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | v6_waybel_1(sK15, sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f3302, f637])).
fof(f3302, plain, (~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | v6_waybel_1(sK15, sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(resolution, [], [f772, f1864])).
fof(f772, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | v6_waybel_1(X1, X0) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f310])).
fof(f310, plain, ! [X0] : (! [X1] : ((v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0))) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f309])).
fof(f309, plain, ! [X0] : (! [X1] : (((v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1)) | (~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0))) | (~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f220])).
fof(f220, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) => ((v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1))))), inference(pure_predicate_removal, [], [f48])).
fof(f48, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) => ((v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v6_waybel_1(X1, X0) & v1_partfun1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', cc3_waybel_1)).
fof(f4604, plain, (~ v6_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | spl60_85), inference(subsumption_resolution, [], [f4586, f4179])).
fof(f4179, plain, (~ sP4(sK15, sK14) | spl60_85), inference(resolution, [], [f4076, f909])).
fof(f909, plain, ! [X0, X1] : (v3_lattice3(k2_yellow_2(X1, X1, X0)) | ~ sP4(X0, X1)), inference(cnf_transformation, [], [f550])).
fof(f550, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X1, X1, X0)) & v2_lattice3(k2_yellow_2(X1, X1, X0)) & v1_lattice3(k2_yellow_2(X1, X1, X0)) & v25_waybel_0(k2_yellow_2(X1, X1, X0)) & v24_waybel_0(k2_yellow_2(X1, X1, X0)) & v4_yellow_0(k2_yellow_2(X1, X1, X0), X1) & v3_yellow_0(k2_yellow_2(X1, X1, X0)) & v2_yellow_0(k2_yellow_2(X1, X1, X0)) & v1_yellow_0(k2_yellow_2(X1, X1, X0)) & v4_orders_2(k2_yellow_2(X1, X1, X0)) & v3_orders_2(k2_yellow_2(X1, X1, X0)) & v2_orders_2(k2_yellow_2(X1, X1, X0)) & v1_orders_2(k2_yellow_2(X1, X1, X0)) & ~ v3_struct_0(k2_yellow_2(X1, X1, X0))) | ~ sP4(X0, X1)), inference(rectify, [], [f549])).
fof(f549, plain, ! [X1, X0] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ sP4(X1, X0)), inference(nnf_transformation, [], [f496])).
fof(f496, plain, ! [X1, X0] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ sP4(X1, X0)), inference(usedef, [], [e496])).
fof(e496, plain, ! [X1, X0] : (sP4(X1, X0) <=> (v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f4076, plain, (~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | spl60_85), inference(avatar_component_clause, [], [f4074])).
fof(f4586, plain, (sP4(sK15, sK14) | ~ v6_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v3_lattice3(sK14) | ~ v2_lattice3(sK14) | ~ v1_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14)), inference(resolution, [], [f910, f1864])).
fof(f910, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | sP4(X1, X0) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f497])).
fof(f497, plain, ! [X0, X1] : (sP4(X1, X0) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(definition_folding, [], [f382, e496])).
fof(f382, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f381])).
fof(f381, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f105])).
fof(f105, plain, ! [X0, X1] : ((m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', fc10_waybel34)).
fof(f4519, plain, (spl60_88 | ~ spl60_11 | spl60_26), inference(avatar_split_clause, [], [f4518, f1652, f1328, f4093])).
fof(f4518, plain, (v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4517, f629])).
fof(f4517, plain, (v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v3_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4516, f630])).
fof(f4516, plain, (v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4515, f632])).
fof(f4515, plain, (v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v2_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4514, f634])).
fof(f4514, plain, (v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | ~ v2_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4513, f3830])).
fof(f4513, plain, (v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ l1_orders_2(sK14) | ~ v2_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4512, f3526])).
fof(f3526, plain, (~ v3_struct_0(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3525, f1653])).
fof(f3525, plain, (~ v3_struct_0(k2_yellow_2(sK14, sK14, sK15)) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3524, f634])).
fof(f3524, plain, (~ v3_struct_0(k2_yellow_2(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3523, f1329])).
fof(f3523, plain, (~ v3_struct_0(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f3522, f636])).
fof(f3522, plain, (~ v3_struct_0(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(duplicate_literal_removal, [], [f3505])).
fof(f3505, plain, (~ v3_struct_0(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(resolution, [], [f953, f1864])).
fof(f953, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v3_struct_0(k2_yellow_2(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f399])).
fof(f399, plain, ! [X0, X1, X2] : ((v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f398])).
fof(f398, plain, ! [X0, X1, X2] : ((v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f114])).
fof(f114, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', fc2_yellow_2)).
fof(f4512, plain, (v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | v3_struct_0(k2_yellow_2(sK14, sK14, sK15)) | ~ m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ l1_orders_2(sK14) | ~ v2_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4509, f3794])).
fof(f4509, plain, (v2_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ v4_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | v3_struct_0(k2_yellow_2(sK14, sK14, sK15)) | ~ m1_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ l1_orders_2(sK14) | ~ v2_lattice3(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(resolution, [], [f4442, f661])).
fof(f661, plain, ! [X0, X1] : (~ v5_yellow_0(X1, X0) | v2_lattice3(X1) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0)), inference(cnf_transformation, [], [f253])).
fof(f253, plain, ! [X0] : (! [X1] : ((v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v2_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1)) | ~ v5_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0)), inference(flattening, [], [f252])).
fof(f252, plain, ! [X0] : (! [X1] : (((v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v2_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1)) | (~ v5_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1))) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ! [X0] : ((l1_orders_2(X0) & v2_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => ((v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & ~ v3_struct_0(X1)) => (v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v2_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', cc11_yellow_0)).
fof(f4442, plain, (v5_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4441, f1653])).
fof(f4441, plain, (v5_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f4440, f628])).
fof(f4440, plain, (v5_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v2_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f4439, f629])).
fof(f4439, plain, (v5_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f4438, f630])).
fof(f4438, plain, (v5_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f4437, f634])).
fof(f4437, plain, (v5_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ l1_orders_2(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f4436, f1329])).
fof(f4436, plain, (v5_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f4435, f636])).
fof(f4435, plain, (v5_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f4417, f637])).
fof(f4417, plain, (v5_yellow_0(k2_yellow_2(sK14, sK14, sK15), sK14) | ~ v7_waybel_1(sK15, sK14) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | ~ v4_orders_2(sK14) | ~ v3_orders_2(sK14) | ~ v2_orders_2(sK14) | v3_struct_0(sK14)), inference(resolution, [], [f894, f1864])).
fof(f894, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f380])).
fof(f380, plain, ! [X0, X1] : ((v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f379])).
fof(f379, plain, ! [X0, X1] : ((v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f199])).
fof(f199, plain, ! [X0, X1] : ((m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1) & l1_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), inference(pure_predicate_removal, [], [f104])).
fof(f104, plain, ! [X0, X1] : ((m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1) & l1_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v3_waybel_0(k2_yellow_2(X0, X0, X1), X0) & v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', fc10_waybel10)).
fof(f4090, plain, (spl60_87 | ~ spl60_85 | ~ spl60_11 | spl60_26 | ~ spl60_66), inference(avatar_split_clause, [], [f4085, f3461, f1652, f1328, f4074, f4087])).
fof(f4085, plain, (~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26 | ~ spl60_66)), inference(subsumption_resolution, [], [f4057, f3526])).
fof(f4057, plain, (~ v3_lattice3(k2_yellow_2(sK14, sK14, sK15)) | v3_struct_0(k2_yellow_2(sK14, sK14, sK15)) | v1_lattice3(k2_yellow_2(sK14, sK14, sK15)) | ~ spl60_66), inference(resolution, [], [f3462, f723])).
fof(f723, plain, ! [X0] : (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | v3_struct_0(X0) | v1_lattice3(X0)), inference(cnf_transformation, [], [f282])).
fof(f282, plain, ! [X0] : ((v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)) | ~ v3_lattice3(X0) | v3_struct_0(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f281])).
fof(f281, plain, ! [X0] : (((v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)) | (~ v3_lattice3(X0) | v3_struct_0(X0))) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : (l1_orders_2(X0) => ((v3_lattice3(X0) & ~ v3_struct_0(X0)) => (v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', cc1_yellow_0)).
fof(f4024, plain, (spl60_66 | ~ spl60_11 | spl60_26), inference(avatar_split_clause, [], [f4023, f1652, f1328, f3461])).
fof(f4023, plain, (l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f4019, f634])).
fof(f4019, plain, (l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | (~ spl60_11 | spl60_26)), inference(resolution, [], [f3830, f875])).
fof(f875, plain, ! [X0, X1] : (~ m1_yellow_0(X1, X0) | l1_orders_2(X1) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f373])).
fof(f373, plain, ! [X0] : (! [X1] : (l1_orders_2(X1) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f90])).
fof(f90, plain, ! [X0] : (l1_orders_2(X0) => ! [X1] : (m1_yellow_0(X1, X0) => l1_orders_2(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', dt_m1_yellow_0)).
fof(f3687, plain, (spl60_70 | ~ spl60_11 | spl60_26), inference(avatar_split_clause, [], [f3686, f1652, f1328, f3657])).
fof(f3686, plain, (v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3685, f1653])).
fof(f3685, plain, (v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3684, f634])).
fof(f3684, plain, (v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3683, f1329])).
fof(f3683, plain, (v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f3682, f636])).
fof(f3682, plain, (v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(duplicate_literal_removal, [], [f3665])).
fof(f3665, plain, (v2_funct_1(k3_waybel_1(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(resolution, [], [f1025, f1864])).
fof(f1025, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v2_funct_1(k3_waybel_1(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f420])).
fof(f420, plain, ! [X0, X1, X2] : ((v5_orders_3(k3_waybel_1(X0, X1, X2), k2_yellow_2(X0, X1, X2), X1) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & ~ v1_xboole_0(k3_waybel_1(X0, X1, X2)) & v2_funct_1(k3_waybel_1(X0, X1, X2)) & v1_funct_1(k3_waybel_1(X0, X1, X2)) & v1_relat_1(k3_waybel_1(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f419])).
fof(f419, plain, ! [X0, X1, X2] : ((v5_orders_3(k3_waybel_1(X0, X1, X2), k2_yellow_2(X0, X1, X2), X1) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & ~ v1_xboole_0(k3_waybel_1(X0, X1, X2)) & v2_funct_1(k3_waybel_1(X0, X1, X2)) & v1_funct_1(k3_waybel_1(X0, X1, X2)) & v1_relat_1(k3_waybel_1(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f221])).
fof(f221, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (v5_orders_3(k3_waybel_1(X0, X1, X2), k2_yellow_2(X0, X1, X2), X1) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & ~ v1_xboole_0(k3_waybel_1(X0, X1, X2)) & v2_funct_1(k3_waybel_1(X0, X1, X2)) & v1_funct_1(k3_waybel_1(X0, X1, X2)) & v1_relat_1(k3_waybel_1(X0, X1, X2)))), inference(pure_predicate_removal, [], [f126])).
fof(f126, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (v5_orders_3(k3_waybel_1(X0, X1, X2), k2_yellow_2(X0, X1, X2), X1) & v1_partfun1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & ~ v1_xboole_0(k3_waybel_1(X0, X1, X2)) & v2_funct_1(k3_waybel_1(X0, X1, X2)) & v1_funct_1(k3_waybel_1(X0, X1, X2)) & v1_relat_1(k3_waybel_1(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', fc7_waybel_1)).
fof(f3468, plain, (~ spl60_66 | spl60_67 | ~ spl60_11 | spl60_26), inference(avatar_split_clause, [], [f3459, f1652, f1328, f3465, f3461])).
fof(f3459, plain, ((k2_yellow_2(sK14, sK14, sK15) = g1_orders_2(u1_struct_0(k2_yellow_2(sK14, sK14, sK15)), u1_orders_2(k2_yellow_2(sK14, sK14, sK15)))) | ~ l1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(resolution, [], [f3444, f649])).
fof(f649, plain, ! [X0] : (~ v1_orders_2(X0) | (g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = X0) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f246])).
fof(f246, plain, ! [X0] : ((g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = X0) | ~ v1_orders_2(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f245])).
fof(f245, plain, ! [X0] : (((g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = X0) | ~ v1_orders_2(X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (l1_orders_2(X0) => (v1_orders_2(X0) => (g1_orders_2(u1_struct_0(X0), u1_orders_2(X0)) = X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', abstractness_v1_orders_2)).
fof(f3444, plain, (v1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | (~ spl60_11 | spl60_26)), inference(subsumption_resolution, [], [f3443, f1653])).
fof(f3443, plain, (v1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3442, f634])).
fof(f3442, plain, (v1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ spl60_11), inference(subsumption_resolution, [], [f3441, f1329])).
fof(f3441, plain, (v1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(subsumption_resolution, [], [f3440, f636])).
fof(f3440, plain, (v1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(duplicate_literal_removal, [], [f3423])).
fof(f3423, plain, (v1_orders_2(k2_yellow_2(sK14, sK14, sK15)) | ~ v1_funct_2(sK15, u1_struct_0(sK14), u1_struct_0(sK14)) | ~ v1_funct_1(sK15) | ~ l1_orders_2(sK14) | v3_struct_0(sK14) | ~ l1_orders_2(sK14) | v3_struct_0(sK14)), inference(resolution, [], [f861, f1864])).
fof(f861, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v1_orders_2(k2_yellow_2(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f365])).
fof(f1753, plain, ~ spl60_26, inference(avatar_contradiction_clause, [], [f1752])).
fof(f1752, plain, ($false | ~ spl60_26), inference(subsumption_resolution, [], [f1751, f634])).
fof(f1751, plain, (~ l1_orders_2(sK14) | ~ spl60_26), inference(subsumption_resolution, [], [f1747, f631])).
fof(f1747, plain, (~ v1_lattice3(sK14) | ~ l1_orders_2(sK14) | ~ spl60_26), inference(resolution, [], [f1654, f701])).
fof(f701, plain, ! [X0] : (~ v3_struct_0(X0) | ~ v1_lattice3(X0) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f268])).
fof(f268, plain, ! [X0] : (~ v3_struct_0(X0) | ~ v1_lattice3(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f267])).
fof(f267, plain, ! [X0] : ((~ v3_struct_0(X0) | ~ v1_lattice3(X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ! [X0] : (l1_orders_2(X0) => (v1_lattice3(X0) => ~ v3_struct_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT368+1.p', cc1_lattice3)).
fof(f1654, plain, (v3_struct_0(sK14) | ~ spl60_26), inference(avatar_component_clause, [], [f1652])).
fof(f1335, plain, spl60_11, inference(avatar_split_clause, [], [f635, f1328])).
fof(f635, plain, v1_funct_1(sK15), inference(cnf_transformation, [], [f524])).
fof(f1334, plain, (spl60_1 | ~ spl60_11 | spl60_12), inference(avatar_split_clause, [], [f1275, f1332, f1328, f1279])).
fof(f1275, plain, ! [X6, X5] : (v1_t_0topsp(sK15, X6, X5) | ~ m2_relset_1(sK15, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_2(sK15, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_1(sK15) | ~ m1_yellow_9(X6, sK14) | ~ v4_waybel11(X6) | ~ m1_yellow_9(X5, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X5) | v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)), inference(equality_resolution, [], [f639])).
fof(f639, plain, ! [X6, X7, X5] : (v1_t_0topsp(X7, X6, X5) | ~ (sK15 = X7) | ~ m2_relset_1(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_2(X7, u1_struct_0(X6), u1_struct_0(X5)) | ~ v1_funct_1(X7) | ~ m1_yellow_9(X6, sK14) | ~ v4_waybel11(X6) | ~ m1_yellow_9(X5, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel11(X5) | v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)), inference(cnf_transformation, [], [f524])).
fof(f1326, plain, (~ spl60_1 | spl60_10), inference(avatar_split_clause, [], [f640, f1323, f1279])).
fof(f640, plain, (v4_waybel11(sK16) | ~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)), inference(cnf_transformation, [], [f524])).
fof(f1321, plain, (~ spl60_1 | spl60_9), inference(avatar_split_clause, [], [f641, f1318, f1279])).
fof(f641, plain, (m1_yellow_9(sK16, k2_yellow_2(sK14, sK14, sK15)) | ~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)), inference(cnf_transformation, [], [f524])).
fof(f1316, plain, (~ spl60_1 | spl60_8), inference(avatar_split_clause, [], [f642, f1313, f1279])).
fof(f642, plain, (v4_waybel11(sK17) | ~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)), inference(cnf_transformation, [], [f524])).
fof(f1311, plain, (~ spl60_1 | spl60_7), inference(avatar_split_clause, [], [f643, f1308, f1279])).
fof(f643, plain, (m1_yellow_9(sK17, sK14) | ~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)), inference(cnf_transformation, [], [f524])).
fof(f1291, plain, (~ spl60_1 | spl60_3), inference(avatar_split_clause, [], [f647, f1288, f1279])).
fof(f647, plain, ((sK15 = sK18) | ~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)), inference(cnf_transformation, [], [f524])).
fof(f1286, plain, (~ spl60_1 | ~ spl60_2), inference(avatar_split_clause, [], [f648, f1283, f1279])).
fof(f648, plain, (~ v1_t_0topsp(sK18, sK17, sK16) | ~ v4_waybel_0(k2_yellow_2(sK14, sK14, sK15), sK14)), inference(cnf_transformation, [], [f524])).