fof(f6406, plain, $false, inference(avatar_sat_refutation, [], [f1627, f1654, f1672, f1715, f1729, f1795, f1821, f1829, f1872, f1895, f2678, f2689, f2694, f2757, f5915, f5918, f5986, f6020, f6134, f6310, f6359, f6380, f6402])).
fof(f6402, plain, ~ spl36_205, inference(avatar_contradiction_clause, [], [f6401])).
fof(f6401, plain, ($false | ~ spl36_205), inference(subsumption_resolution, [], [f6392, f390])).
fof(f390, plain, ~ v2_setfam_1(sK5), inference(cnf_transformation, [], [f321])).
fof(f321, plain, (~ r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), k9_waybel34(sK5), k8_waybel34(sK5)) & ~ v2_setfam_1(sK5)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5])], [f177, f320])).
fof(f320, plain, (? [X0] : (~ r3_yellow20(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0), k9_waybel34(X0), k8_waybel34(X0)) & ~ v2_setfam_1(X0)) => (~ r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), k9_waybel34(sK5), k8_waybel34(sK5)) & ~ v2_setfam_1(sK5))), introduced(choice_axiom, [])).
fof(f177, plain, ? [X0] : (~ r3_yellow20(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0), k9_waybel34(X0), k8_waybel34(X0)) & ~ v2_setfam_1(X0)), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : (~ v2_setfam_1(X0) => r3_yellow20(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0), k9_waybel34(X0), k8_waybel34(X0))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : (~ v2_setfam_1(X0) => r3_yellow20(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0), k9_waybel34(X0), k8_waybel34(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', t57_waybel34)).
fof(f6392, plain, (v2_setfam_1(sK5) | ~ spl36_205), inference(resolution, [], [f6384, f444])).
fof(f444, plain, ! [X0] : (~ v1_xboole_0(X0) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ! [X0] : (~ v1_xboole_0(X0) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ! [X0] : (~ v2_setfam_1(X0) => ~ v1_xboole_0(X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', cc2_setfam_1)).
fof(f6384, plain, (v1_xboole_0(sK5) | ~ spl36_205), inference(resolution, [], [f5914, f525])).
fof(f525, plain, ! [X0] : (~ v3_struct_0(k8_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ! [X0] : ((m1_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) & v3_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) & v6_altcat_1(k8_waybel34(X0)) & v2_altcat_1(k8_waybel34(X0)) & ~ v3_struct_0(k8_waybel34(X0))) | v1_xboole_0(X0)), inference(ennf_transformation, [], [f55])).
fof(f55, plain, ! [X0] : (~ v1_xboole_0(X0) => (m1_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) & v3_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) & v6_altcat_1(k8_waybel34(X0)) & v2_altcat_1(k8_waybel34(X0)) & ~ v3_struct_0(k8_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', dt_k8_waybel34)).
fof(f5914, plain, (v3_struct_0(k8_waybel34(sK5)) | ~ spl36_205), inference(avatar_component_clause, [], [f5912])).
fof(f5912, plain, (spl36_205 <=> v3_struct_0(k8_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_205])])).
fof(f6380, plain, spl36_203, inference(avatar_contradiction_clause, [], [f6379])).
fof(f6379, plain, ($false | spl36_203), inference(subsumption_resolution, [], [f6370, f390])).
fof(f6370, plain, (v2_setfam_1(sK5) | spl36_203), inference(resolution, [], [f6341, f444])).
fof(f6341, plain, (v1_xboole_0(sK5) | spl36_203), inference(resolution, [], [f5906, f528])).
fof(f528, plain, ! [X0] : (v3_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f235])).
fof(f5906, plain, (~ v3_altcat_2(k8_waybel34(sK5), k4_waybel34(sK5)) | spl36_203), inference(avatar_component_clause, [], [f5904])).
fof(f5904, plain, (spl36_203 <=> v3_altcat_2(k8_waybel34(sK5), k4_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_203])])).
fof(f6359, plain, spl36_204, inference(avatar_contradiction_clause, [], [f6358])).
fof(f6358, plain, ($false | spl36_204), inference(subsumption_resolution, [], [f6349, f390])).
fof(f6349, plain, (v2_setfam_1(sK5) | spl36_204), inference(resolution, [], [f6340, f444])).
fof(f6340, plain, (v1_xboole_0(sK5) | spl36_204), inference(resolution, [], [f5910, f526])).
fof(f526, plain, ! [X0] : (v2_altcat_1(k8_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f235])).
fof(f5910, plain, (~ v2_altcat_1(k8_waybel34(sK5)) | spl36_204), inference(avatar_component_clause, [], [f5908])).
fof(f5908, plain, (spl36_204 <=> v2_altcat_1(k8_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_204])])).
fof(f6310, plain, spl36_202, inference(avatar_contradiction_clause, [], [f6309])).
fof(f6309, plain, ($false | spl36_202), inference(subsumption_resolution, [], [f6300, f390])).
fof(f6300, plain, (v2_setfam_1(sK5) | spl36_202), inference(resolution, [], [f6292, f444])).
fof(f6292, plain, (v1_xboole_0(sK5) | spl36_202), inference(resolution, [], [f5902, f529])).
fof(f529, plain, ! [X0] : (m1_altcat_2(k8_waybel34(X0), k4_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f235])).
fof(f5902, plain, (~ m1_altcat_2(k8_waybel34(sK5), k4_waybel34(sK5)) | spl36_202), inference(avatar_component_clause, [], [f5900])).
fof(f5900, plain, (spl36_202 <=> m1_altcat_2(k8_waybel34(sK5), k4_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_202])])).
fof(f6134, plain, spl36_199, inference(avatar_contradiction_clause, [], [f6133])).
fof(f6133, plain, ($false | spl36_199), inference(subsumption_resolution, [], [f6132, f390])).
fof(f6132, plain, (v2_setfam_1(sK5) | spl36_199), inference(resolution, [], [f5890, f533])).
fof(f533, plain, ! [X0] : (v3_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f236])).
fof(f236, plain, ! [X0] : ((m1_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) & v3_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) & v6_altcat_1(k9_waybel34(X0)) & v2_altcat_1(k9_waybel34(X0)) & ~ v3_struct_0(k9_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f56])).
fof(f56, plain, ! [X0] : (~ v2_setfam_1(X0) => (m1_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) & v3_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) & v6_altcat_1(k9_waybel34(X0)) & v2_altcat_1(k9_waybel34(X0)) & ~ v3_struct_0(k9_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', dt_k9_waybel34)).
fof(f5890, plain, (~ v3_altcat_2(k9_waybel34(sK5), k5_waybel34(sK5)) | spl36_199), inference(avatar_component_clause, [], [f5888])).
fof(f5888, plain, (spl36_199 <=> v3_altcat_2(k9_waybel34(sK5), k5_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_199])])).
fof(f6020, plain, ~ spl36_201, inference(avatar_contradiction_clause, [], [f6019])).
fof(f6019, plain, ($false | ~ spl36_201), inference(subsumption_resolution, [], [f6018, f390])).
fof(f6018, plain, (v2_setfam_1(sK5) | ~ spl36_201), inference(resolution, [], [f5898, f530])).
fof(f530, plain, ! [X0] : (~ v3_struct_0(k9_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f236])).
fof(f5898, plain, (v3_struct_0(k9_waybel34(sK5)) | ~ spl36_201), inference(avatar_component_clause, [], [f5896])).
fof(f5896, plain, (spl36_201 <=> v3_struct_0(k9_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_201])])).
fof(f5986, plain, spl36_200, inference(avatar_contradiction_clause, [], [f5985])).
fof(f5985, plain, ($false | spl36_200), inference(subsumption_resolution, [], [f5984, f390])).
fof(f5984, plain, (v2_setfam_1(sK5) | spl36_200), inference(resolution, [], [f5894, f531])).
fof(f531, plain, ! [X0] : (v2_altcat_1(k9_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f236])).
fof(f5894, plain, (~ v2_altcat_1(k9_waybel34(sK5)) | spl36_200), inference(avatar_component_clause, [], [f5892])).
fof(f5892, plain, (spl36_200 <=> v2_altcat_1(k9_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_200])])).
fof(f5918, plain, spl36_198, inference(avatar_contradiction_clause, [], [f5917])).
fof(f5917, plain, ($false | spl36_198), inference(subsumption_resolution, [], [f5916, f390])).
fof(f5916, plain, (v2_setfam_1(sK5) | spl36_198), inference(resolution, [], [f5886, f534])).
fof(f534, plain, ! [X0] : (m1_altcat_2(k9_waybel34(X0), k5_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f236])).
fof(f5886, plain, (~ m1_altcat_2(k9_waybel34(sK5), k5_waybel34(sK5)) | spl36_198), inference(avatar_component_clause, [], [f5884])).
fof(f5884, plain, (spl36_198 <=> m1_altcat_2(k9_waybel34(sK5), k5_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_198])])).
fof(f5915, plain, (~ spl36_198 | ~ spl36_199 | ~ spl36_200 | spl36_201 | ~ spl36_202 | ~ spl36_203 | ~ spl36_204 | spl36_205 | ~ spl36_22), inference(avatar_split_clause, [], [f5882, f1595, f5912, f5908, f5904, f5900, f5896, f5892, f5888, f5884])).
fof(f1595, plain, (spl36_22 <=> ! [X1, X0] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | v3_struct_0(X1) | ~ v2_altcat_1(X1) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | v3_struct_0(X0) | ~ v2_altcat_1(X0) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl36_22])])).
fof(f5882, plain, (v3_struct_0(k8_waybel34(sK5)) | ~ v2_altcat_1(k8_waybel34(sK5)) | ~ v3_altcat_2(k8_waybel34(sK5), k4_waybel34(sK5)) | ~ m1_altcat_2(k8_waybel34(sK5), k4_waybel34(sK5)) | v3_struct_0(k9_waybel34(sK5)) | ~ v2_altcat_1(k9_waybel34(sK5)) | ~ v3_altcat_2(k9_waybel34(sK5), k5_waybel34(sK5)) | ~ m1_altcat_2(k9_waybel34(sK5), k5_waybel34(sK5)) | ~ spl36_22), inference(subsumption_resolution, [], [f5881, f390])).
fof(f5881, plain, (v3_struct_0(k8_waybel34(sK5)) | ~ v2_altcat_1(k8_waybel34(sK5)) | ~ v3_altcat_2(k8_waybel34(sK5), k4_waybel34(sK5)) | ~ m1_altcat_2(k8_waybel34(sK5), k4_waybel34(sK5)) | v3_struct_0(k9_waybel34(sK5)) | ~ v2_altcat_1(k9_waybel34(sK5)) | ~ v3_altcat_2(k9_waybel34(sK5), k5_waybel34(sK5)) | ~ m1_altcat_2(k9_waybel34(sK5), k5_waybel34(sK5)) | v2_setfam_1(sK5) | ~ spl36_22), inference(subsumption_resolution, [], [f5880, f391])).
fof(f391, plain, ~ r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), k9_waybel34(sK5), k8_waybel34(sK5)), inference(cnf_transformation, [], [f321])).
fof(f5880, plain, (v3_struct_0(k8_waybel34(sK5)) | ~ v2_altcat_1(k8_waybel34(sK5)) | ~ v3_altcat_2(k8_waybel34(sK5), k4_waybel34(sK5)) | ~ m1_altcat_2(k8_waybel34(sK5), k4_waybel34(sK5)) | v3_struct_0(k9_waybel34(sK5)) | ~ v2_altcat_1(k9_waybel34(sK5)) | ~ v3_altcat_2(k9_waybel34(sK5), k5_waybel34(sK5)) | ~ m1_altcat_2(k9_waybel34(sK5), k5_waybel34(sK5)) | r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), k9_waybel34(sK5), k8_waybel34(sK5)) | v2_setfam_1(sK5) | ~ spl36_22), inference(resolution, [], [f1596, f679])).
fof(f679, plain, ! [X0] : (r3_yellow20(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0), k8_waybel34(X0), k9_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f305])).
fof(f305, plain, ! [X0] : (r3_yellow20(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0), k8_waybel34(X0), k9_waybel34(X0)) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f124])).
fof(f124, plain, ! [X0] : (~ v2_setfam_1(X0) => r3_yellow20(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0), k8_waybel34(X0), k9_waybel34(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', t56_waybel34)).
fof(f1596, plain, (! [X0, X1] : (~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | v3_struct_0(X1) | ~ v2_altcat_1(X1) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | v3_struct_0(X0) | ~ v2_altcat_1(X0) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1)) | ~ spl36_22), inference(avatar_component_clause, [], [f1595])).
fof(f2757, plain, spl36_21, inference(avatar_contradiction_clause, [], [f2756])).
fof(f2756, plain, ($false | spl36_21), inference(subsumption_resolution, [], [f2755, f390])).
fof(f2755, plain, (v2_setfam_1(sK5) | spl36_21), inference(resolution, [], [f1593, f605])).
fof(f605, plain, ! [X0] : (v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f277])).
fof(f277, plain, ! [X0] : ((v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v14_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v11_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f152])).
fof(f152, plain, ! [X0] : (~ v2_setfam_1(X0) => (v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v14_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v11_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)))), inference(pure_predicate_removal, [], [f143])).
fof(f143, plain, ! [X0] : (~ v2_setfam_1(X0) => (v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v14_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v11_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v8_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)))), inference(pure_predicate_removal, [], [f135])).
fof(f135, plain, ! [X0] : (~ v2_setfam_1(X0) => (v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v14_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v12_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v11_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v8_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)))), inference(pure_predicate_removal, [], [f95])).
fof(f95, plain, ! [X0] : (~ v2_setfam_1(X0) => (v21_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v14_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v12_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v11_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v8_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v6_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', fc7_waybel34)).
fof(f1593, plain, (~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | spl36_21), inference(avatar_component_clause, [], [f1591])).
fof(f1591, plain, (spl36_21 <=> v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_21])])).
fof(f2694, plain, spl36_20, inference(avatar_contradiction_clause, [], [f2693])).
fof(f2693, plain, ($false | spl36_20), inference(subsumption_resolution, [], [f2692, f390])).
fof(f2692, plain, (v2_setfam_1(sK5) | spl36_20), inference(resolution, [], [f1589, f521])).
fof(f521, plain, ! [X0] : (m2_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f233])).
fof(f233, plain, ! [X0] : ((m2_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f53])).
fof(f53, plain, ! [X0] : (~ v2_setfam_1(X0) => (m2_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) & v9_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', dt_k6_waybel34)).
fof(f1589, plain, (~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | spl36_20), inference(avatar_component_clause, [], [f1587])).
fof(f1587, plain, (spl36_20 <=> m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_20])])).
fof(f2689, plain, (~ spl36_19 | ~ spl36_20 | ~ spl36_21 | spl36_22 | spl36_9 | ~ spl36_10 | ~ spl36_11 | ~ spl36_12 | ~ spl36_13 | spl36_14 | ~ spl36_15 | ~ spl36_16 | ~ spl36_17 | ~ spl36_18), inference(avatar_split_clause, [], [f2688, f1579, f1575, f1571, f1567, f1563, f1559, f1555, f1551, f1547, f1543, f1595, f1591, f1587, f1583])).
fof(f1583, plain, (spl36_19 <=> v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_19])])).
fof(f1543, plain, (spl36_9 <=> v3_struct_0(k4_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_9])])).
fof(f1547, plain, (spl36_10 <=> v2_altcat_1(k4_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_10])])).
fof(f1551, plain, (spl36_11 <=> v11_altcat_1(k4_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_11])])).
fof(f1555, plain, (spl36_12 <=> v12_altcat_1(k4_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_12])])).
fof(f1559, plain, (spl36_13 <=> l2_altcat_1(k4_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_13])])).
fof(f1563, plain, (spl36_14 <=> v3_struct_0(k5_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_14])])).
fof(f1567, plain, (spl36_15 <=> v2_altcat_1(k5_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_15])])).
fof(f1571, plain, (spl36_16 <=> v11_altcat_1(k5_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_16])])).
fof(f1575, plain, (spl36_17 <=> v12_altcat_1(k5_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_17])])).
fof(f1579, plain, (spl36_18 <=> l2_altcat_1(k5_waybel34(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl36_18])])).
fof(f2688, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5))) | (spl36_9 | ~ spl36_10 | ~ spl36_11 | ~ spl36_12 | ~ spl36_13 | spl36_14 | ~ spl36_15 | ~ spl36_16 | ~ spl36_17 | ~ spl36_18)), inference(subsumption_resolution, [], [f2687, f1544])).
fof(f1544, plain, (~ v3_struct_0(k4_waybel34(sK5)) | spl36_9), inference(avatar_component_clause, [], [f1543])).
fof(f2687, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | v3_struct_0(k4_waybel34(sK5))) | (~ spl36_10 | ~ spl36_11 | ~ spl36_12 | ~ spl36_13 | spl36_14 | ~ spl36_15 | ~ spl36_16 | ~ spl36_17 | ~ spl36_18)), inference(subsumption_resolution, [], [f2686, f1548])).
fof(f1548, plain, (v2_altcat_1(k4_waybel34(sK5)) | ~ spl36_10), inference(avatar_component_clause, [], [f1547])).
fof(f2686, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v2_altcat_1(k4_waybel34(sK5)) | v3_struct_0(k4_waybel34(sK5))) | (~ spl36_11 | ~ spl36_12 | ~ spl36_13 | spl36_14 | ~ spl36_15 | ~ spl36_16 | ~ spl36_17 | ~ spl36_18)), inference(subsumption_resolution, [], [f2685, f1552])).
fof(f1552, plain, (v11_altcat_1(k4_waybel34(sK5)) | ~ spl36_11), inference(avatar_component_clause, [], [f1551])).
fof(f2685, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v11_altcat_1(k4_waybel34(sK5)) | ~ v2_altcat_1(k4_waybel34(sK5)) | v3_struct_0(k4_waybel34(sK5))) | (~ spl36_12 | ~ spl36_13 | spl36_14 | ~ spl36_15 | ~ spl36_16 | ~ spl36_17 | ~ spl36_18)), inference(subsumption_resolution, [], [f2684, f1556])).
fof(f1556, plain, (v12_altcat_1(k4_waybel34(sK5)) | ~ spl36_12), inference(avatar_component_clause, [], [f1555])).
fof(f2684, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v12_altcat_1(k4_waybel34(sK5)) | ~ v11_altcat_1(k4_waybel34(sK5)) | ~ v2_altcat_1(k4_waybel34(sK5)) | v3_struct_0(k4_waybel34(sK5))) | (~ spl36_13 | spl36_14 | ~ spl36_15 | ~ spl36_16 | ~ spl36_17 | ~ spl36_18)), inference(subsumption_resolution, [], [f2683, f1560])).
fof(f1560, plain, (l2_altcat_1(k4_waybel34(sK5)) | ~ spl36_13), inference(avatar_component_clause, [], [f1559])).
fof(f2683, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ l2_altcat_1(k4_waybel34(sK5)) | ~ v12_altcat_1(k4_waybel34(sK5)) | ~ v11_altcat_1(k4_waybel34(sK5)) | ~ v2_altcat_1(k4_waybel34(sK5)) | v3_struct_0(k4_waybel34(sK5))) | (spl36_14 | ~ spl36_15 | ~ spl36_16 | ~ spl36_17 | ~ spl36_18)), inference(subsumption_resolution, [], [f2682, f1564])).
fof(f1564, plain, (~ v3_struct_0(k5_waybel34(sK5)) | spl36_14), inference(avatar_component_clause, [], [f1563])).
fof(f2682, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | v3_struct_0(k5_waybel34(sK5)) | ~ l2_altcat_1(k4_waybel34(sK5)) | ~ v12_altcat_1(k4_waybel34(sK5)) | ~ v11_altcat_1(k4_waybel34(sK5)) | ~ v2_altcat_1(k4_waybel34(sK5)) | v3_struct_0(k4_waybel34(sK5))) | (~ spl36_15 | ~ spl36_16 | ~ spl36_17 | ~ spl36_18)), inference(subsumption_resolution, [], [f2681, f1568])).
fof(f1568, plain, (v2_altcat_1(k5_waybel34(sK5)) | ~ spl36_15), inference(avatar_component_clause, [], [f1567])).
fof(f2681, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v2_altcat_1(k5_waybel34(sK5)) | v3_struct_0(k5_waybel34(sK5)) | ~ l2_altcat_1(k4_waybel34(sK5)) | ~ v12_altcat_1(k4_waybel34(sK5)) | ~ v11_altcat_1(k4_waybel34(sK5)) | ~ v2_altcat_1(k4_waybel34(sK5)) | v3_struct_0(k4_waybel34(sK5))) | (~ spl36_16 | ~ spl36_17 | ~ spl36_18)), inference(subsumption_resolution, [], [f2680, f1572])).
fof(f1572, plain, (v11_altcat_1(k5_waybel34(sK5)) | ~ spl36_16), inference(avatar_component_clause, [], [f1571])).
fof(f2680, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v11_altcat_1(k5_waybel34(sK5)) | ~ v2_altcat_1(k5_waybel34(sK5)) | v3_struct_0(k5_waybel34(sK5)) | ~ l2_altcat_1(k4_waybel34(sK5)) | ~ v12_altcat_1(k4_waybel34(sK5)) | ~ v11_altcat_1(k4_waybel34(sK5)) | ~ v2_altcat_1(k4_waybel34(sK5)) | v3_struct_0(k4_waybel34(sK5))) | (~ spl36_17 | ~ spl36_18)), inference(subsumption_resolution, [], [f2679, f1576])).
fof(f1576, plain, (v12_altcat_1(k5_waybel34(sK5)) | ~ spl36_17), inference(avatar_component_clause, [], [f1575])).
fof(f2679, plain, (! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v12_altcat_1(k5_waybel34(sK5)) | ~ v11_altcat_1(k5_waybel34(sK5)) | ~ v2_altcat_1(k5_waybel34(sK5)) | v3_struct_0(k5_waybel34(sK5)) | ~ l2_altcat_1(k4_waybel34(sK5)) | ~ v12_altcat_1(k4_waybel34(sK5)) | ~ v11_altcat_1(k4_waybel34(sK5)) | ~ v2_altcat_1(k4_waybel34(sK5)) | v3_struct_0(k4_waybel34(sK5))) | ~ spl36_18), inference(subsumption_resolution, [], [f1882, f1580])).
fof(f1580, plain, (l2_altcat_1(k5_waybel34(sK5)) | ~ spl36_18), inference(avatar_component_clause, [], [f1579])).
fof(f1882, plain, ! [X0, X1] : (r3_yellow20(k5_waybel34(sK5), k4_waybel34(sK5), k7_waybel34(sK5), X0, X1) | ~ r3_yellow20(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5), X1, X0) | ~ m1_altcat_2(X0, k5_waybel34(sK5)) | ~ v3_altcat_2(X0, k5_waybel34(sK5)) | ~ v2_altcat_1(X0) | v3_struct_0(X0) | ~ m1_altcat_2(X1, k4_waybel34(sK5)) | ~ v3_altcat_2(X1, k4_waybel34(sK5)) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ v21_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ m2_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | ~ l2_altcat_1(k5_waybel34(sK5)) | ~ v12_altcat_1(k5_waybel34(sK5)) | ~ v11_altcat_1(k5_waybel34(sK5)) | ~ v2_altcat_1(k5_waybel34(sK5)) | v3_struct_0(k5_waybel34(sK5)) | ~ l2_altcat_1(k4_waybel34(sK5)) | ~ v12_altcat_1(k4_waybel34(sK5)) | ~ v11_altcat_1(k4_waybel34(sK5)) | ~ v2_altcat_1(k4_waybel34(sK5)) | v3_struct_0(k4_waybel34(sK5))), inference(superposition, [], [f678, f942])).
fof(f942, plain, (k7_waybel34(sK5) = k15_functor0(k4_waybel34(sK5), k5_waybel34(sK5), k6_waybel34(sK5))), inference(resolution, [], [f672, f390])).
fof(f672, plain, ! [X0] : (v2_setfam_1(X0) | (k7_waybel34(X0) = k15_functor0(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0)))), inference(cnf_transformation, [], [f296])).
fof(f296, plain, ! [X0] : (((k6_waybel34(X0) = k15_functor0(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0))) & (k7_waybel34(X0) = k15_functor0(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0)))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f118])).
fof(f118, plain, ! [X0] : (~ v2_setfam_1(X0) => ((k6_waybel34(X0) = k15_functor0(k5_waybel34(X0), k4_waybel34(X0), k7_waybel34(X0))) & (k7_waybel34(X0) = k15_functor0(k4_waybel34(X0), k5_waybel34(X0), k6_waybel34(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', t18_waybel34)).
fof(f678, plain, ! [X4, X2, X0, X3, X1] : (r3_yellow20(X1, X0, k15_functor0(X0, X1, X2), X4, X3) | ~ r3_yellow20(X0, X1, X2, X3, X4) | ~ m1_altcat_2(X4, X1) | ~ v3_altcat_2(X4, X1) | ~ v2_altcat_1(X4) | v3_struct_0(X4) | ~ m1_altcat_2(X3, X0) | ~ v3_altcat_2(X3, X0) | ~ v2_altcat_1(X3) | v3_struct_0(X3) | ~ v21_functor0(X2, X0, X1) | ~ m2_functor0(X2, X0, X1) | ~ v16_functor0(X2, X0, X1) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1) | ~ l2_altcat_1(X0) | ~ v12_altcat_1(X0) | ~ v11_altcat_1(X0) | ~ v2_altcat_1(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f304])).
fof(f304, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (r3_yellow20(X1, X0, k15_functor0(X0, X1, X2), X4, X3) | ~ r3_yellow20(X0, X1, X2, X3, X4) | ~ m1_altcat_2(X4, X1) | ~ v3_altcat_2(X4, X1) | ~ v2_altcat_1(X4) | v3_struct_0(X4)) | ~ m1_altcat_2(X3, X0) | ~ v3_altcat_2(X3, X0) | ~ v2_altcat_1(X3) | v3_struct_0(X3)) | ~ v21_functor0(X2, X0, X1) | ~ m2_functor0(X2, X0, X1) | ~ v16_functor0(X2, X0, X1)) | ~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1)) | ~ l2_altcat_1(X0) | ~ v12_altcat_1(X0) | ~ v11_altcat_1(X0) | ~ v2_altcat_1(X0) | v3_struct_0(X0)), inference(flattening, [], [f303])).
fof(f303, plain, ! [X0] : (! [X1] : (! [X2] : ((! [X3] : (! [X4] : ((r3_yellow20(X1, X0, k15_functor0(X0, X1, X2), X4, X3) | ~ r3_yellow20(X0, X1, X2, X3, X4)) | (~ m1_altcat_2(X4, X1) | ~ v3_altcat_2(X4, X1) | ~ v2_altcat_1(X4) | v3_struct_0(X4))) | (~ m1_altcat_2(X3, X0) | ~ v3_altcat_2(X3, X0) | ~ v2_altcat_1(X3) | v3_struct_0(X3))) | ~ v21_functor0(X2, X0, X1)) | (~ m2_functor0(X2, X0, X1) | ~ v16_functor0(X2, X0, X1))) | (~ l2_altcat_1(X1) | ~ v12_altcat_1(X1) | ~ v11_altcat_1(X1) | ~ v2_altcat_1(X1) | v3_struct_0(X1))) | (~ l2_altcat_1(X0) | ~ v12_altcat_1(X0) | ~ v11_altcat_1(X0) | ~ v2_altcat_1(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f123])).
fof(f123, plain, ! [X0] : ((l2_altcat_1(X0) & v12_altcat_1(X0) & v11_altcat_1(X0) & v2_altcat_1(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((l2_altcat_1(X1) & v12_altcat_1(X1) & v11_altcat_1(X1) & v2_altcat_1(X1) & ~ v3_struct_0(X1)) => ! [X2] : ((m2_functor0(X2, X0, X1) & v16_functor0(X2, X0, X1)) => (v21_functor0(X2, X0, X1) => ! [X3] : ((m1_altcat_2(X3, X0) & v3_altcat_2(X3, X0) & v2_altcat_1(X3) & ~ v3_struct_0(X3)) => ! [X4] : ((m1_altcat_2(X4, X1) & v3_altcat_2(X4, X1) & v2_altcat_1(X4) & ~ v3_struct_0(X4)) => (r3_yellow20(X0, X1, X2, X3, X4) => r3_yellow20(X1, X0, k15_functor0(X0, X1, X2), X4, X3)))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', t52_yellow20)).
fof(f2678, plain, spl36_19, inference(avatar_contradiction_clause, [], [f2677])).
fof(f2677, plain, ($false | spl36_19), inference(subsumption_resolution, [], [f2676, f390])).
fof(f2676, plain, (v2_setfam_1(sK5) | spl36_19), inference(resolution, [], [f1585, f520])).
fof(f520, plain, ! [X0] : (v16_functor0(k6_waybel34(X0), k4_waybel34(X0), k5_waybel34(X0)) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f233])).
fof(f1585, plain, (~ v16_functor0(k6_waybel34(sK5), k4_waybel34(sK5), k5_waybel34(sK5)) | spl36_19), inference(avatar_component_clause, [], [f1583])).
fof(f1895, plain, ~ spl36_14, inference(avatar_contradiction_clause, [], [f1894])).
fof(f1894, plain, ($false | ~ spl36_14), inference(subsumption_resolution, [], [f1893, f390])).
fof(f1893, plain, (v2_setfam_1(sK5) | ~ spl36_14), inference(resolution, [], [f1891, f598])).
fof(f598, plain, ! [X0] : (sP4(X0) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f319])).
fof(f319, plain, ! [X0] : (sP4(X0) | v2_setfam_1(X0)), inference(definition_folding, [], [f274, e318])).
fof(f318, plain, ! [X0] : ((v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0))) | ~ sP4(X0)), inference(usedef, [], [e318])).
fof(e318, plain, ! [X0] : (sP4(X0) <=> (v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f274, plain, ! [X0] : ((v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f165])).
fof(f165, plain, ! [X0] : (~ v2_setfam_1(X0) => (v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0)))), inference(pure_predicate_removal, [], [f93])).
fof(f93, plain, ! [X0] : (~ v2_setfam_1(X0) => (v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v4_yellow18(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', fc6_waybel34)).
fof(f1891, plain, (~ sP4(sK5) | ~ spl36_14), inference(resolution, [], [f1565, f586])).
fof(f586, plain, ! [X0] : (~ v3_struct_0(k5_waybel34(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f354])).
fof(f354, plain, ! [X0] : ((v3_yellow21(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v1_yellow21(k5_waybel34(X0)) & v3_yellow18(k5_waybel34(X0)) & v2_yellow18(k5_waybel34(X0)) & v1_altcat_2(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v9_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0))) | ~ sP4(X0)), inference(nnf_transformation, [], [f318])).
fof(f1565, plain, (v3_struct_0(k5_waybel34(sK5)) | ~ spl36_14), inference(avatar_component_clause, [], [f1563])).
fof(f1872, plain, spl36_18, inference(avatar_contradiction_clause, [], [f1871])).
fof(f1871, plain, ($false | spl36_18), inference(subsumption_resolution, [], [f1845, f704])).
fof(f704, plain, v2_setfam_1(k1_xboole_0), inference(resolution, [], [f444, f580])).
fof(f580, plain, v1_xboole_0(k1_xboole_0), inference(cnf_transformation, [], [f92])).
fof(f92, plain, (v5_membered(k1_xboole_0) & v4_membered(k1_xboole_0) & v3_membered(k1_xboole_0) & v2_membered(k1_xboole_0) & v1_membered(k1_xboole_0) & v1_xboole_0(k1_xboole_0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', fc6_membered)).
fof(f1845, plain, (~ v2_setfam_1(k1_xboole_0) | spl36_18), inference(backward_demodulation, [], [f390, f1840])).
fof(f1840, plain, ((k1_xboole_0 = sK5) | spl36_18), inference(resolution, [], [f1831, f681])).
fof(f681, plain, ! [X0] : (~ v1_xboole_0(X0) | (k1_xboole_0 = X0)), inference(cnf_transformation, [], [f307])).
fof(f307, plain, ! [X0] : ((k1_xboole_0 = X0) | ~ v1_xboole_0(X0)), inference(ennf_transformation, [], [f126])).
fof(f126, plain, ! [X0] : (v1_xboole_0(X0) => (k1_xboole_0 = X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', t6_boole)).
fof(f1831, plain, (v1_xboole_0(sK5) | spl36_18), inference(resolution, [], [f1581, f518])).
fof(f518, plain, ! [X0] : (l2_altcat_1(k5_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f232])).
fof(f232, plain, ! [X0] : ((l2_altcat_1(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0))) | v1_xboole_0(X0)), inference(ennf_transformation, [], [f52])).
fof(f52, plain, ! [X0] : (~ v1_xboole_0(X0) => (l2_altcat_1(k5_waybel34(X0)) & v2_yellow21(k5_waybel34(X0)) & v12_altcat_1(k5_waybel34(X0)) & v11_altcat_1(k5_waybel34(X0)) & v6_altcat_1(k5_waybel34(X0)) & v2_altcat_1(k5_waybel34(X0)) & ~ v3_struct_0(k5_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', dt_k5_waybel34)).
fof(f1581, plain, (~ l2_altcat_1(k5_waybel34(sK5)) | spl36_18), inference(avatar_component_clause, [], [f1579])).
fof(f1829, plain, spl36_17, inference(avatar_contradiction_clause, [], [f1828])).
fof(f1828, plain, ($false | spl36_17), inference(subsumption_resolution, [], [f1827, f390])).
fof(f1827, plain, (v2_setfam_1(sK5) | spl36_17), inference(resolution, [], [f1825, f598])).
fof(f1825, plain, (~ sP4(sK5) | spl36_17), inference(resolution, [], [f1577, f591])).
fof(f591, plain, ! [X0] : (v12_altcat_1(k5_waybel34(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f354])).
fof(f1577, plain, (~ v12_altcat_1(k5_waybel34(sK5)) | spl36_17), inference(avatar_component_clause, [], [f1575])).
fof(f1821, plain, spl36_16, inference(avatar_contradiction_clause, [], [f1820])).
fof(f1820, plain, ($false | spl36_16), inference(subsumption_resolution, [], [f1819, f390])).
fof(f1819, plain, (v2_setfam_1(sK5) | spl36_16), inference(resolution, [], [f1814, f598])).
fof(f1814, plain, (~ sP4(sK5) | spl36_16), inference(resolution, [], [f1573, f590])).
fof(f590, plain, ! [X0] : (v11_altcat_1(k5_waybel34(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f354])).
fof(f1573, plain, (~ v11_altcat_1(k5_waybel34(sK5)) | spl36_16), inference(avatar_component_clause, [], [f1571])).
fof(f1795, plain, spl36_15, inference(avatar_contradiction_clause, [], [f1794])).
fof(f1794, plain, ($false | spl36_15), inference(subsumption_resolution, [], [f1793, f390])).
fof(f1793, plain, (v2_setfam_1(sK5) | spl36_15), inference(resolution, [], [f1788, f598])).
fof(f1788, plain, (~ sP4(sK5) | spl36_15), inference(resolution, [], [f1569, f587])).
fof(f587, plain, ! [X0] : (v2_altcat_1(k5_waybel34(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f354])).
fof(f1569, plain, (~ v2_altcat_1(k5_waybel34(sK5)) | spl36_15), inference(avatar_component_clause, [], [f1567])).
fof(f1729, plain, ~ spl36_9, inference(avatar_contradiction_clause, [], [f1728])).
fof(f1728, plain, ($false | ~ spl36_9), inference(subsumption_resolution, [], [f1727, f390])).
fof(f1727, plain, (v2_setfam_1(sK5) | ~ spl36_9), inference(resolution, [], [f1725, f579])).
fof(f579, plain, ! [X0] : (sP3(X0) | v2_setfam_1(X0)), inference(cnf_transformation, [], [f317])).
fof(f317, plain, ! [X0] : (sP3(X0) | v2_setfam_1(X0)), inference(definition_folding, [], [f273, e316])).
fof(f316, plain, ! [X0] : ((v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0))) | ~ sP3(X0)), inference(usedef, [], [e316])).
fof(e316, plain, ! [X0] : (sP3(X0) <=> (v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f273, plain, ! [X0] : ((v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0))) | v2_setfam_1(X0)), inference(ennf_transformation, [], [f166])).
fof(f166, plain, ! [X0] : (~ v2_setfam_1(X0) => (v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0)))), inference(pure_predicate_removal, [], [f91])).
fof(f91, plain, ! [X0] : (~ v2_setfam_1(X0) => (v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v4_yellow18(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', fc5_waybel34)).
fof(f1725, plain, (~ sP3(sK5) | ~ spl36_9), inference(resolution, [], [f1545, f567])).
fof(f567, plain, ! [X0] : (~ v3_struct_0(k4_waybel34(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f353])).
fof(f353, plain, ! [X0] : ((v3_yellow21(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v1_yellow21(k4_waybel34(X0)) & v3_yellow18(k4_waybel34(X0)) & v2_yellow18(k4_waybel34(X0)) & v1_altcat_2(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v9_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0))) | ~ sP3(X0)), inference(nnf_transformation, [], [f316])).
fof(f1545, plain, (v3_struct_0(k4_waybel34(sK5)) | ~ spl36_9), inference(avatar_component_clause, [], [f1543])).
fof(f1715, plain, spl36_13, inference(avatar_contradiction_clause, [], [f1714])).
fof(f1714, plain, ($false | spl36_13), inference(subsumption_resolution, [], [f1697, f704])).
fof(f1697, plain, (~ v2_setfam_1(k1_xboole_0) | spl36_13), inference(backward_demodulation, [], [f390, f1692])).
fof(f1692, plain, ((k1_xboole_0 = sK5) | spl36_13), inference(resolution, [], [f1680, f681])).
fof(f1680, plain, (v1_xboole_0(sK5) | spl36_13), inference(resolution, [], [f1561, f511])).
fof(f511, plain, ! [X0] : (l2_altcat_1(k4_waybel34(X0)) | v1_xboole_0(X0)), inference(cnf_transformation, [], [f231])).
fof(f231, plain, ! [X0] : ((l2_altcat_1(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0))) | v1_xboole_0(X0)), inference(ennf_transformation, [], [f51])).
fof(f51, plain, ! [X0] : (~ v1_xboole_0(X0) => (l2_altcat_1(k4_waybel34(X0)) & v2_yellow21(k4_waybel34(X0)) & v12_altcat_1(k4_waybel34(X0)) & v11_altcat_1(k4_waybel34(X0)) & v6_altcat_1(k4_waybel34(X0)) & v2_altcat_1(k4_waybel34(X0)) & ~ v3_struct_0(k4_waybel34(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT376+1.p', dt_k4_waybel34)).
fof(f1561, plain, (~ l2_altcat_1(k4_waybel34(sK5)) | spl36_13), inference(avatar_component_clause, [], [f1559])).
fof(f1672, plain, spl36_12, inference(avatar_contradiction_clause, [], [f1671])).
fof(f1671, plain, ($false | spl36_12), inference(subsumption_resolution, [], [f1670, f390])).
fof(f1670, plain, (v2_setfam_1(sK5) | spl36_12), inference(resolution, [], [f1665, f579])).
fof(f1665, plain, (~ sP3(sK5) | spl36_12), inference(resolution, [], [f1557, f572])).
fof(f572, plain, ! [X0] : (v12_altcat_1(k4_waybel34(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f353])).
fof(f1557, plain, (~ v12_altcat_1(k4_waybel34(sK5)) | spl36_12), inference(avatar_component_clause, [], [f1555])).
fof(f1654, plain, spl36_10, inference(avatar_contradiction_clause, [], [f1653])).
fof(f1653, plain, ($false | spl36_10), inference(subsumption_resolution, [], [f1652, f390])).
fof(f1652, plain, (v2_setfam_1(sK5) | spl36_10), inference(resolution, [], [f1647, f579])).
fof(f1647, plain, (~ sP3(sK5) | spl36_10), inference(resolution, [], [f1549, f568])).
fof(f568, plain, ! [X0] : (v2_altcat_1(k4_waybel34(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f353])).
fof(f1549, plain, (~ v2_altcat_1(k4_waybel34(sK5)) | spl36_10), inference(avatar_component_clause, [], [f1547])).
fof(f1627, plain, spl36_11, inference(avatar_contradiction_clause, [], [f1626])).
fof(f1626, plain, ($false | spl36_11), inference(subsumption_resolution, [], [f1625, f390])).
fof(f1625, plain, (v2_setfam_1(sK5) | spl36_11), inference(resolution, [], [f1620, f579])).
fof(f1620, plain, (~ sP3(sK5) | spl36_11), inference(resolution, [], [f1553, f571])).
fof(f571, plain, ! [X0] : (v11_altcat_1(k4_waybel34(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f353])).
fof(f1553, plain, (~ v11_altcat_1(k4_waybel34(sK5)) | spl36_11), inference(avatar_component_clause, [], [f1551])).