fof(f372784, plain, $false, inference(avatar_sat_refutation, [], [f823, f916, f926, f1265, f1268, f2679, f2702, f3235, f3241, f3244, f3258, f3268, f20806, f25074, f372595, f372672, f372783])).
fof(f372783, plain, (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_68 | ~ spl29_136 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084), inference(avatar_contradiction_clause, [], [f372782])).
fof(f372782, plain, ($false | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_68 | ~ spl29_136 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084)), inference(subsumption_resolution, [], [f372781, f291])).
fof(f291, plain, m1_subset_1(sK2, u1_struct_0(sK1)), inference(cnf_transformation, [], [f230])).
fof(f230, plain, ((! [X2] : (~ r2_filter_2(sK1, X2) | ~ r2_hidden(sK2, X2) | ~ m2_filter_2(X2, sK1)) & ~ (k6_lattices(sK1) = sK2) & v14_lattices(sK1) & m1_subset_1(sK2, u1_struct_0(sK1))) & l3_lattices(sK1) & v10_lattices(sK1) & ~ v3_struct_0(sK1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2])], [f111, f229, f228])).
fof(f228, plain, (? [X0] : (? [X1] : (! [X2] : (~ r2_filter_2(X0, X2) | ~ r2_hidden(X1, X2) | ~ m2_filter_2(X2, X0)) & ~ (k6_lattices(X0) = X1) & v14_lattices(X0) & m1_subset_1(X1, u1_struct_0(X0))) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (? [X1] : (! [X2] : (~ r2_filter_2(sK1, X2) | ~ r2_hidden(X1, X2) | ~ m2_filter_2(X2, sK1)) & ~ (k6_lattices(sK1) = X1) & v14_lattices(sK1) & m1_subset_1(X1, u1_struct_0(sK1))) & l3_lattices(sK1) & v10_lattices(sK1) & ~ v3_struct_0(sK1))), introduced(choice_axiom, [])).
fof(f229, plain, (? [X1] : (! [X2] : (~ r2_filter_2(sK1, X2) | ~ r2_hidden(X1, X2) | ~ m2_filter_2(X2, sK1)) & ~ (k6_lattices(sK1) = X1) & v14_lattices(sK1) & m1_subset_1(X1, u1_struct_0(sK1))) => (! [X2] : (~ r2_filter_2(sK1, X2) | ~ r2_hidden(sK2, X2) | ~ m2_filter_2(X2, sK1)) & ~ (k6_lattices(sK1) = sK2) & v14_lattices(sK1) & m1_subset_1(sK2, u1_struct_0(sK1)))), introduced(choice_axiom, [])).
fof(f111, plain, ? [X0] : (? [X1] : (! [X2] : (~ r2_filter_2(X0, X2) | ~ r2_hidden(X1, X2) | ~ m2_filter_2(X2, X0)) & ~ (k6_lattices(X0) = X1) & v14_lattices(X0) & m1_subset_1(X1, u1_struct_0(X0))) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f110])).
fof(f110, plain, ? [X0] : (? [X1] : ((! [X2] : ((~ r2_filter_2(X0, X2) | ~ r2_hidden(X1, X2)) | ~ m2_filter_2(X2, X0)) & ~ (k6_lattices(X0) = X1) & v14_lattices(X0)) & m1_subset_1(X1, u1_struct_0(X0))) & (l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ~ (! [X2] : (m2_filter_2(X2, X0) => ~ (r2_filter_2(X0, X2) & r2_hidden(X1, X2))) & ~ (k6_lattices(X0) = X1) & v14_lattices(X0)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ~ (! [X2] : (m2_filter_2(X2, X0) => ~ (r2_filter_2(X0, X2) & r2_hidden(X1, X2))) & ~ (k6_lattices(X0) = X1) & v14_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', t36_filter_2)).
fof(f372781, plain, (~ m1_subset_1(sK2, u1_struct_0(sK1)) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_68 | ~ spl29_136 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084)), inference(forward_demodulation, [], [f372780, f4861])).
fof(f4861, plain, ((u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | ~ spl29_68), inference(trivial_inequality_removal, [], [f4844])).
fof(f4844, plain, (~ (k1_lattice2(sK1) = k1_lattice2(sK1)) | (u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | ~ spl29_68), inference(superposition, [], [f1264, f1712])).
fof(f1712, plain, (k1_lattice2(sK1) = g3_lattices(u1_struct_0(k1_lattice2(sK1)), u2_lattices(k1_lattice2(sK1)), u1_lattices(k1_lattice2(sK1)))), inference(resolution, [], [f745, f290])).
fof(f290, plain, l3_lattices(sK1), inference(cnf_transformation, [], [f230])).
fof(f745, plain, ! [X1] : (~ l3_lattices(X1) | (k1_lattice2(X1) = g3_lattices(u1_struct_0(k1_lattice2(X1)), u2_lattices(k1_lattice2(X1)), u1_lattices(k1_lattice2(X1))))), inference(subsumption_resolution, [], [f740, f324])).
fof(f324, plain, ! [X0] : (l3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ! [X0] : ((l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ! [X0] : (l3_lattices(X0) => (l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', dt_k1_lattice2)).
fof(f740, plain, ! [X1] : ((k1_lattice2(X1) = g3_lattices(u1_struct_0(k1_lattice2(X1)), u2_lattices(k1_lattice2(X1)), u1_lattices(k1_lattice2(X1)))) | ~ l3_lattices(k1_lattice2(X1)) | ~ l3_lattices(X1)), inference(resolution, [], [f295, f323])).
fof(f323, plain, ! [X0] : (v3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f140])).
fof(f295, plain, ! [X0] : (~ v3_lattices(X0) | (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ! [X0] : ((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f112])).
fof(f112, plain, ! [X0] : (((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (l3_lattices(X0) => (v3_lattices(X0) => (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', abstractness_v3_lattices)).
fof(f1264, plain, (! [X17, X15, X16] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15)) | ~ spl29_68), inference(avatar_component_clause, [], [f1263])).
fof(f1263, plain, (spl29_68 <=> ! [X16, X15, X17] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15))), introduced(avatar_definition, [new_symbols(naming, [spl29_68])])).
fof(f372780, plain, (~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_136 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084)), inference(subsumption_resolution, [], [f372779, f293])).
fof(f293, plain, ~ (k6_lattices(sK1) = sK2), inference(cnf_transformation, [], [f230])).
fof(f372779, plain, ((k6_lattices(sK1) = sK2) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_136 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084)), inference(forward_demodulation, [], [f372778, f780])).
fof(f780, plain, (k6_lattices(sK1) = k5_lattices(k1_lattice2(sK1))), inference(subsumption_resolution, [], [f779, f288])).
fof(f288, plain, ~ v3_struct_0(sK1), inference(cnf_transformation, [], [f230])).
fof(f779, plain, ((k6_lattices(sK1) = k5_lattices(k1_lattice2(sK1))) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f778, f289])).
fof(f289, plain, v10_lattices(sK1), inference(cnf_transformation, [], [f230])).
fof(f778, plain, ((k6_lattices(sK1) = k5_lattices(k1_lattice2(sK1))) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f776, f290])).
fof(f776, plain, (~ l3_lattices(sK1) | (k6_lattices(sK1) = k5_lattices(k1_lattice2(sK1))) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f461, f292])).
fof(f292, plain, v14_lattices(sK1), inference(cnf_transformation, [], [f230])).
fof(f461, plain, ! [X0] : (~ v14_lattices(X0) | ~ l3_lattices(X0) | (k6_lattices(X0) = k5_lattices(k1_lattice2(X0))) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f223])).
fof(f223, plain, ! [X0] : ((k6_lattices(X0) = k5_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v14_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f222])).
fof(f222, plain, ! [X0] : ((k6_lattices(X0) = k5_lattices(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v14_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f95])).
fof(f95, plain, ! [X0] : ((l3_lattices(X0) & v14_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k6_lattices(X0) = k5_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', t79_lattice2)).
fof(f372778, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_136 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084)), inference(subsumption_resolution, [], [f372777, f787])).
fof(f787, plain, (~ v3_struct_0(k1_lattice2(sK1)) | spl29_13), inference(avatar_component_clause, [], [f786])).
fof(f786, plain, (spl29_13 <=> v3_struct_0(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_13])])).
fof(f372777, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (~ spl29_16 | ~ spl29_17 | ~ spl29_136 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084)), inference(subsumption_resolution, [], [f372776, f879])).
fof(f879, plain, (v10_lattices(k1_lattice2(sK1)) | ~ spl29_16), inference(avatar_component_clause, [], [f878])).
fof(f878, plain, (spl29_16 <=> v10_lattices(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_16])])).
fof(f372776, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (~ spl29_16 | ~ spl29_17 | ~ spl29_136 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084)), inference(subsumption_resolution, [], [f372775, f2662])).
fof(f2662, plain, (v13_lattices(k1_lattice2(sK1)) | ~ spl29_136), inference(avatar_component_clause, [], [f2661])).
fof(f2661, plain, (spl29_136 <=> v13_lattices(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_136])])).
fof(f372775, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (~ spl29_16 | ~ spl29_17 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084)), inference(subsumption_resolution, [], [f372774, f883])).
fof(f883, plain, (l3_lattices(k1_lattice2(sK1)) | ~ spl29_17), inference(avatar_component_clause, [], [f882])).
fof(f882, plain, (spl29_17 <=> l3_lattices(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_17])])).
fof(f372774, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (~ spl29_16 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084)), inference(resolution, [], [f372702, f465])).
fof(f465, plain, ! [X0, X1] : (r2_hidden(X1, sK28(X0, X1)) | (k5_lattices(X0) = X1) | ~ l3_lattices(X0) | ~ v13_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0))), inference(duplicate_literal_removal, [], [f450])).
fof(f450, plain, ! [X0, X1] : (r2_hidden(X1, sK28(X0, X1)) | (k5_lattices(X0) = X1) | ~ l3_lattices(X0) | ~ v13_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f285])).
fof(f285, plain, ! [X0] : (! [X1] : ((v1_filter_0(sK28(X0, X1), X0) & r2_hidden(X1, sK28(X0, X1)) & m1_filter_0(sK28(X0, X1), X0)) | (k5_lattices(X0) = X1) | ~ l3_lattices(X0) | ~ v13_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK28])], [f210, f284])).
fof(f284, plain, ! [X1, X0] : (? [X2] : (v1_filter_0(X2, X0) & r2_hidden(X1, X2) & m1_filter_0(X2, X0)) => (v1_filter_0(sK28(X0, X1), X0) & r2_hidden(X1, sK28(X0, X1)) & m1_filter_0(sK28(X0, X1), X0))), introduced(choice_axiom, [])).
fof(f210, plain, ! [X0] : (! [X1] : (? [X2] : (v1_filter_0(X2, X0) & r2_hidden(X1, X2) & m1_filter_0(X2, X0)) | (k5_lattices(X0) = X1) | ~ l3_lattices(X0) | ~ v13_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f209])).
fof(f209, plain, ! [X0] : (! [X1] : ((? [X2] : ((v1_filter_0(X2, X0) & r2_hidden(X1, X2)) & m1_filter_0(X2, X0)) | (k5_lattices(X0) = X1) | ~ l3_lattices(X0) | ~ v13_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f87])).
fof(f87, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ~ (! [X2] : (m1_filter_0(X2, X0) => ~ (v1_filter_0(X2, X0) & r2_hidden(X1, X2))) & ~ (k5_lattices(X0) = X1) & l3_lattices(X0) & v13_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', t24_filter_0)).
fof(f372702, plain, (~ r2_hidden(sK2, sK28(k1_lattice2(sK1), sK2)) | (~ spl29_16 | ~ spl29_882 | ~ spl29_885 | ~ spl29_5084)), inference(subsumption_resolution, [], [f372701, f25198])).
fof(f25198, plain, (m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | (~ spl29_16 | ~ spl29_882 | ~ spl29_885)), inference(backward_demodulation, [], [f20805, f25194])).
fof(f25194, plain, ((sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_16 | ~ spl29_882)), inference(subsumption_resolution, [], [f25193, f879])).
fof(f25193, plain, ((sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_882), inference(subsumption_resolution, [], [f25192, f288])).
fof(f25192, plain, (v3_struct_0(sK1) | (sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_882), inference(subsumption_resolution, [], [f25191, f289])).
fof(f25191, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | (sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_882), inference(subsumption_resolution, [], [f25172, f290])).
fof(f25172, plain, (~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_882), inference(resolution, [], [f20785, f861])).
fof(f861, plain, ! [X2, X3] : (~ m1_filter_0(X3, k1_lattice2(X2)) | ~ l3_lattices(X2) | ~ v10_lattices(X2) | v3_struct_0(X2) | (k8_filter_2(X2, X3) = X3) | ~ v10_lattices(k1_lattice2(X2))), inference(subsumption_resolution, [], [f860, f360])).
fof(f360, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ! [X0] : ((v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f175])).
fof(f175, plain, ! [X0] : ((v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : ((l3_lattices(X0) & ~ v3_struct_0(X0)) => (v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', fc1_lattice2)).
fof(f860, plain, ! [X2, X3] : ((k8_filter_2(X2, X3) = X3) | ~ l3_lattices(X2) | ~ v10_lattices(X2) | v3_struct_0(X2) | ~ m1_filter_0(X3, k1_lattice2(X2)) | ~ v10_lattices(k1_lattice2(X2)) | v3_struct_0(k1_lattice2(X2))), inference(subsumption_resolution, [], [f851, f324])).
fof(f851, plain, ! [X2, X3] : ((k8_filter_2(X2, X3) = X3) | ~ l3_lattices(X2) | ~ v10_lattices(X2) | v3_struct_0(X2) | ~ m1_filter_0(X3, k1_lattice2(X2)) | ~ l3_lattices(k1_lattice2(X2)) | ~ v10_lattices(k1_lattice2(X2)) | v3_struct_0(k1_lattice2(X2))), inference(resolution, [], [f318, f335])).
fof(f335, plain, ! [X0, X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ! [X0] : (! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f154])).
fof(f154, plain, ! [X0] : (! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) | ~ m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_0(X1, X0) => (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', dt_m1_filter_0)).
fof(f318, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) | (k8_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ! [X0] : (! [X1] : ((k8_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0))))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f132])).
fof(f132, plain, ! [X0] : (! [X1] : ((k8_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0))))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) => (k8_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', d7_filter_2)).
fof(f20785, plain, (m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | ~ spl29_882), inference(avatar_component_clause, [], [f20784])).
fof(f20784, plain, (spl29_882 <=> m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_882])])).
fof(f20805, plain, (m2_filter_2(k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)), sK1) | ~ spl29_885), inference(avatar_component_clause, [], [f20803])).
fof(f20803, plain, (spl29_885 <=> m2_filter_2(k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)), sK1)), introduced(avatar_definition, [new_symbols(naming, [spl29_885])])).
fof(f372701, plain, (~ r2_hidden(sK2, sK28(k1_lattice2(sK1), sK2)) | ~ m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | ~ spl29_5084), inference(resolution, [], [f372584, f294])).
fof(f294, plain, ! [X2] : (~ r2_filter_2(sK1, X2) | ~ r2_hidden(sK2, X2) | ~ m2_filter_2(X2, sK1)), inference(cnf_transformation, [], [f230])).
fof(f372584, plain, (r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | ~ spl29_5084), inference(avatar_component_clause, [], [f372583])).
fof(f372583, plain, (spl29_5084 <=> r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))), introduced(avatar_definition, [new_symbols(naming, [spl29_5084])])).
fof(f372672, plain, (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_68 | ~ spl29_136 | spl29_5085), inference(avatar_contradiction_clause, [], [f372671])).
fof(f372671, plain, ($false | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_68 | ~ spl29_136 | spl29_5085)), inference(subsumption_resolution, [], [f372670, f291])).
fof(f372670, plain, (~ m1_subset_1(sK2, u1_struct_0(sK1)) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_68 | ~ spl29_136 | spl29_5085)), inference(forward_demodulation, [], [f372669, f4861])).
fof(f372669, plain, (~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_136 | spl29_5085)), inference(subsumption_resolution, [], [f372668, f293])).
fof(f372668, plain, ((k6_lattices(sK1) = sK2) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_136 | spl29_5085)), inference(forward_demodulation, [], [f372667, f780])).
fof(f372667, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_136 | spl29_5085)), inference(subsumption_resolution, [], [f372666, f787])).
fof(f372666, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (~ spl29_16 | ~ spl29_17 | ~ spl29_136 | spl29_5085)), inference(subsumption_resolution, [], [f372665, f879])).
fof(f372665, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (~ spl29_17 | ~ spl29_136 | spl29_5085)), inference(subsumption_resolution, [], [f372664, f2662])).
fof(f372664, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (~ spl29_17 | spl29_5085)), inference(subsumption_resolution, [], [f372663, f883])).
fof(f372663, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | spl29_5085), inference(resolution, [], [f372588, f466])).
fof(f466, plain, ! [X0, X1] : (v1_filter_0(sK28(X0, X1), X0) | (k5_lattices(X0) = X1) | ~ l3_lattices(X0) | ~ v13_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0))), inference(duplicate_literal_removal, [], [f451])).
fof(f451, plain, ! [X0, X1] : (v1_filter_0(sK28(X0, X1), X0) | (k5_lattices(X0) = X1) | ~ l3_lattices(X0) | ~ v13_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f285])).
fof(f372588, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | spl29_5085), inference(avatar_component_clause, [], [f372587])).
fof(f372587, plain, (spl29_5085 <=> v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_5085])])).
fof(f372595, plain, (spl29_5084 | ~ spl29_5085 | ~ spl29_16 | ~ spl29_138 | ~ spl29_882 | ~ spl29_885), inference(avatar_split_clause, [], [f372594, f20803, f20784, f2676, f878, f372587, f372583])).
fof(f2676, plain, (spl29_138 <=> (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))), introduced(avatar_definition, [new_symbols(naming, [spl29_138])])).
fof(f372594, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | (~ spl29_16 | ~ spl29_138 | ~ spl29_882 | ~ spl29_885)), inference(subsumption_resolution, [], [f372593, f288])).
fof(f372593, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | v3_struct_0(sK1) | (~ spl29_16 | ~ spl29_138 | ~ spl29_882 | ~ spl29_885)), inference(subsumption_resolution, [], [f372592, f289])).
fof(f372592, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (~ spl29_16 | ~ spl29_138 | ~ spl29_882 | ~ spl29_885)), inference(subsumption_resolution, [], [f372591, f290])).
fof(f372591, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (~ spl29_16 | ~ spl29_138 | ~ spl29_882 | ~ spl29_885)), inference(subsumption_resolution, [], [f372562, f25198])).
fof(f372562, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | ~ m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (~ spl29_16 | ~ spl29_138 | ~ spl29_882)), inference(superposition, [], [f454, f25211])).
fof(f25211, plain, ((sK28(k1_lattice2(sK1), sK2) = k15_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_16 | ~ spl29_138 | ~ spl29_882)), inference(backward_demodulation, [], [f25201, f25210])).
fof(f25210, plain, ((sK28(k1_lattice2(sK1), sK2) = k7_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_16 | ~ spl29_138 | ~ spl29_882)), inference(forward_demodulation, [], [f25209, f25195])).
fof(f25195, plain, ((sK28(k1_lattice2(sK1), sK2) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_16 | ~ spl29_138 | ~ spl29_882)), inference(backward_demodulation, [], [f2678, f25194])).
fof(f2678, plain, ((k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ spl29_138), inference(avatar_component_clause, [], [f2676])).
fof(f25209, plain, ((k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | ~ spl29_882), inference(subsumption_resolution, [], [f25208, f290])).
fof(f25208, plain, (~ l3_lattices(sK1) | (k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | ~ spl29_882), inference(subsumption_resolution, [], [f25207, f288])).
fof(f25207, plain, (v3_struct_0(sK1) | ~ l3_lattices(sK1) | (k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | ~ spl29_882), inference(subsumption_resolution, [], [f25174, f289])).
fof(f25174, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ l3_lattices(sK1) | (k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | ~ spl29_882), inference(resolution, [], [f20785, f1782])).
fof(f1782, plain, ! [X4, X3] : (~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v10_lattices(X3) | v3_struct_0(X3) | ~ l3_lattices(X3) | (k14_filter_2(X3, X4) = k7_filter_2(X3, k14_filter_2(X3, X4)))), inference(duplicate_literal_removal, [], [f1781])).
fof(f1781, plain, ! [X4, X3] : (~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | ~ m1_filter_0(X4, k1_lattice2(X3)) | ~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | (k14_filter_2(X3, X4) = k7_filter_2(X3, k14_filter_2(X3, X4)))), inference(resolution, [], [f816, f843])).
fof(f843, plain, ! [X0, X1] : (~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | (k7_filter_2(X0, X1) = X1)), inference(duplicate_literal_removal, [], [f833])).
fof(f833, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f317, f340])).
fof(f340, plain, ! [X0, X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f160])).
fof(f160, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_lattice4(X1, X0) => m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', dt_m2_lattice4)).
fof(f317, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | (k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f130])).
fof(f130, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) => (k7_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', d6_filter_2)).
fof(f816, plain, ! [X0, X1] : (m2_lattice4(k14_filter_2(X1, X0), X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ m1_filter_0(X0, k1_lattice2(X1))), inference(duplicate_literal_removal, [], [f813])).
fof(f813, plain, ! [X0, X1] : (~ m1_filter_0(X0, k1_lattice2(X1)) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | m2_lattice4(k14_filter_2(X1, X0), X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1)), inference(resolution, [], [f321, f339])).
fof(f339, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f158])).
fof(f158, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', dt_m2_filter_2)).
fof(f321, plain, ! [X0, X1] : (m2_filter_2(k14_filter_2(X0, X1), X0) | ~ m1_filter_0(X1, k1_lattice2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ! [X0, X1] : (m2_filter_2(k14_filter_2(X0, X1), X0) | ~ m1_filter_0(X1, k1_lattice2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f136])).
fof(f136, plain, ! [X0, X1] : (m2_filter_2(k14_filter_2(X0, X1), X0) | (~ m1_filter_0(X1, k1_lattice2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0, X1] : ((m1_filter_0(X1, k1_lattice2(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m2_filter_2(k14_filter_2(X0, X1), X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', dt_k14_filter_2)).
fof(f25201, plain, ((k7_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k15_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_16 | ~ spl29_138 | ~ spl29_882)), inference(backward_demodulation, [], [f25190, f25194])).
fof(f25190, plain, ((k7_filter_2(sK1, k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (~ spl29_138 | ~ spl29_882)), inference(forward_demodulation, [], [f25189, f2678])).
fof(f25189, plain, ((k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | ~ spl29_882), inference(subsumption_resolution, [], [f25188, f288])).
fof(f25188, plain, (v3_struct_0(sK1) | (k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | ~ spl29_882), inference(subsumption_resolution, [], [f25187, f289])).
fof(f25187, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | (k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | ~ spl29_882), inference(subsumption_resolution, [], [f25171, f290])).
fof(f25171, plain, (~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | ~ spl29_882), inference(resolution, [], [f20785, f848])).
fof(f848, plain, ! [X2, X1] : (~ m1_filter_0(X2, k1_lattice2(X1)) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | (k7_filter_2(X1, k14_filter_2(X1, X2)) = k15_filter_2(X1, k14_filter_2(X1, X2)))), inference(duplicate_literal_removal, [], [f847])).
fof(f847, plain, ! [X2, X1] : ((k7_filter_2(X1, k14_filter_2(X1, X2)) = k15_filter_2(X1, k14_filter_2(X1, X2))) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ m1_filter_0(X2, k1_lattice2(X1)) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1)), inference(resolution, [], [f442, f321])).
fof(f442, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | (k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f204])).
fof(f204, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | (~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f82])).
fof(f82, plain, ! [X0, X1] : ((m2_filter_2(X1, X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k7_filter_2(X0, X1) = k15_filter_2(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', redefinition_k15_filter_2)).
fof(f454, plain, ! [X0, X1] : (~ v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | r2_filter_2(X0, X1) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f286])).
fof(f286, plain, ! [X0] : (! [X1] : (((r2_filter_2(X0, X1) | ~ v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) & (v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ r2_filter_2(X0, X1))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f214])).
fof(f214, plain, ! [X0] : (! [X1] : ((r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f213])).
fof(f213, plain, ! [X0] : (! [X1] : ((r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f89])).
fof(f89, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', t33_filter_2)).
fof(f25074, plain, (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_68 | ~ spl29_136 | spl29_882), inference(avatar_contradiction_clause, [], [f25073])).
fof(f25073, plain, ($false | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_68 | ~ spl29_136 | spl29_882)), inference(subsumption_resolution, [], [f25072, f291])).
fof(f25072, plain, (~ m1_subset_1(sK2, u1_struct_0(sK1)) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_68 | ~ spl29_136 | spl29_882)), inference(forward_demodulation, [], [f25071, f4861])).
fof(f25071, plain, (~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_136 | spl29_882)), inference(subsumption_resolution, [], [f25070, f293])).
fof(f25070, plain, ((k6_lattices(sK1) = sK2) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_136 | spl29_882)), inference(forward_demodulation, [], [f25069, f780])).
fof(f25069, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (spl29_13 | ~ spl29_16 | ~ spl29_17 | ~ spl29_136 | spl29_882)), inference(subsumption_resolution, [], [f25068, f787])).
fof(f25068, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (~ spl29_16 | ~ spl29_17 | ~ spl29_136 | spl29_882)), inference(subsumption_resolution, [], [f25067, f879])).
fof(f25067, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (~ spl29_17 | ~ spl29_136 | spl29_882)), inference(subsumption_resolution, [], [f25066, f2662])).
fof(f25066, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | (~ spl29_17 | spl29_882)), inference(subsumption_resolution, [], [f25065, f883])).
fof(f25065, plain, ((sK2 = k5_lattices(k1_lattice2(sK1))) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | spl29_882), inference(resolution, [], [f20786, f464])).
fof(f464, plain, ! [X0, X1] : (m1_filter_0(sK28(X0, X1), X0) | (k5_lattices(X0) = X1) | ~ l3_lattices(X0) | ~ v13_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0))), inference(duplicate_literal_removal, [], [f449])).
fof(f449, plain, ! [X0, X1] : (m1_filter_0(sK28(X0, X1), X0) | (k5_lattices(X0) = X1) | ~ l3_lattices(X0) | ~ v13_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f285])).
fof(f20786, plain, (~ m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | spl29_882), inference(avatar_component_clause, [], [f20784])).
fof(f20806, plain, (~ spl29_882 | spl29_885 | ~ spl29_138), inference(avatar_split_clause, [], [f20801, f2676, f20803, f20784])).
fof(f20801, plain, (m2_filter_2(k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)), sK1) | ~ m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | ~ spl29_138), inference(subsumption_resolution, [], [f20800, f288])).
fof(f20800, plain, (m2_filter_2(k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)), sK1) | ~ m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | v3_struct_0(sK1) | ~ spl29_138), inference(subsumption_resolution, [], [f20799, f289])).
fof(f20799, plain, (m2_filter_2(k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)), sK1) | ~ m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ spl29_138), inference(subsumption_resolution, [], [f20779, f290])).
fof(f20779, plain, (m2_filter_2(k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)), sK1) | ~ m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ spl29_138), inference(superposition, [], [f321, f2678])).
fof(f3268, plain, spl29_65, inference(avatar_contradiction_clause, [], [f3267])).
fof(f3267, plain, ($false | spl29_65), inference(subsumption_resolution, [], [f3266, f480])).
fof(f480, plain, l2_lattices(sK1), inference(resolution, [], [f333, f290])).
fof(f333, plain, ! [X0] : (~ l3_lattices(X0) | l2_lattices(X0)), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ! [X0] : ((l2_lattices(X0) & l1_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', dt_l3_lattices)).
fof(f3266, plain, (~ l2_lattices(sK1) | spl29_65), inference(resolution, [], [f1253, f733])).
fof(f733, plain, ! [X1] : (m1_relset_1(u2_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l2_lattices(X1)), inference(resolution, [], [f347, f445])).
fof(f445, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f283])).
fof(f283, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f84])).
fof(f84, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', redefinition_m2_relset_1)).
fof(f347, plain, ! [X0] : (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f164])).
fof(f164, plain, ! [X0] : ((m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0))) | ~ l2_lattices(X0)), inference(ennf_transformation, [], [f41])).
fof(f41, plain, ! [X0] : (l2_lattices(X0) => (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', dt_u2_lattices)).
fof(f1253, plain, (~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl29_65), inference(avatar_component_clause, [], [f1251])).
fof(f1251, plain, (spl29_65 <=> m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_65])])).
fof(f3258, plain, spl29_64, inference(avatar_contradiction_clause, [], [f3257])).
fof(f3257, plain, ($false | spl29_64), inference(subsumption_resolution, [], [f3256, f480])).
fof(f3256, plain, (~ l2_lattices(sK1) | spl29_64), inference(resolution, [], [f1249, f346])).
fof(f346, plain, ! [X0] : (v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f164])).
fof(f1249, plain, (~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl29_64), inference(avatar_component_clause, [], [f1247])).
fof(f1247, plain, (spl29_64 <=> v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_64])])).
fof(f3244, plain, spl29_63, inference(avatar_contradiction_clause, [], [f3243])).
fof(f3243, plain, ($false | spl29_63), inference(subsumption_resolution, [], [f3242, f480])).
fof(f3242, plain, (~ l2_lattices(sK1) | spl29_63), inference(resolution, [], [f1245, f345])).
fof(f345, plain, ! [X0] : (v1_funct_1(u2_lattices(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f164])).
fof(f1245, plain, (~ v1_funct_1(u2_lattices(sK1)) | spl29_63), inference(avatar_component_clause, [], [f1243])).
fof(f1243, plain, (spl29_63 <=> v1_funct_1(u2_lattices(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_63])])).
fof(f3241, plain, spl29_62, inference(avatar_contradiction_clause, [], [f3240])).
fof(f3240, plain, ($false | spl29_62), inference(subsumption_resolution, [], [f3239, f473])).
fof(f473, plain, l1_lattices(sK1), inference(resolution, [], [f332, f290])).
fof(f332, plain, ! [X0] : (~ l3_lattices(X0) | l1_lattices(X0)), inference(cnf_transformation, [], [f153])).
fof(f3239, plain, (~ l1_lattices(sK1) | spl29_62), inference(resolution, [], [f1241, f731])).
fof(f731, plain, ! [X1] : (m1_relset_1(u1_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l1_lattices(X1)), inference(resolution, [], [f344, f445])).
fof(f344, plain, ! [X0] : (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ! [X0] : ((m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0))) | ~ l1_lattices(X0)), inference(ennf_transformation, [], [f39])).
fof(f39, plain, ! [X0] : (l1_lattices(X0) => (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', dt_u1_lattices)).
fof(f1241, plain, (~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl29_62), inference(avatar_component_clause, [], [f1239])).
fof(f1239, plain, (spl29_62 <=> m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_62])])).
fof(f3235, plain, spl29_61, inference(avatar_contradiction_clause, [], [f3234])).
fof(f3234, plain, ($false | spl29_61), inference(subsumption_resolution, [], [f3233, f473])).
fof(f3233, plain, (~ l1_lattices(sK1) | spl29_61), inference(resolution, [], [f1237, f343])).
fof(f343, plain, ! [X0] : (v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f163])).
fof(f1237, plain, (~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl29_61), inference(avatar_component_clause, [], [f1235])).
fof(f1235, plain, (spl29_61 <=> v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_61])])).
fof(f2702, plain, spl29_136, inference(avatar_contradiction_clause, [], [f2701])).
fof(f2701, plain, ($false | spl29_136), inference(subsumption_resolution, [], [f2700, f288])).
fof(f2700, plain, (v3_struct_0(sK1) | spl29_136), inference(subsumption_resolution, [], [f2699, f289])).
fof(f2699, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | spl29_136), inference(subsumption_resolution, [], [f2698, f290])).
fof(f2698, plain, (~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | spl29_136), inference(subsumption_resolution, [], [f2697, f292])).
fof(f2697, plain, (~ v14_lattices(sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | spl29_136), inference(resolution, [], [f2663, f458])).
fof(f458, plain, ! [X0] : (v13_lattices(k1_lattice2(X0)) | ~ v14_lattices(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f287])).
fof(f287, plain, ! [X0] : (((v14_lattices(X0) | ~ v13_lattices(k1_lattice2(X0))) & (v13_lattices(k1_lattice2(X0)) | ~ v14_lattices(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f220])).
fof(f220, plain, ! [X0] : ((v14_lattices(X0) <=> v13_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f219])).
fof(f219, plain, ! [X0] : ((v14_lattices(X0) <=> v13_lattices(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f93])).
fof(f93, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v14_lattices(X0) <=> v13_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', t64_lattice2)).
fof(f2663, plain, (~ v13_lattices(k1_lattice2(sK1)) | spl29_136), inference(avatar_component_clause, [], [f2661])).
fof(f2679, plain, (spl29_138 | ~ spl29_136 | ~ spl29_16), inference(avatar_split_clause, [], [f2674, f878, f2661, f2676])).
fof(f2674, plain, (~ v13_lattices(k1_lattice2(sK1)) | (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ spl29_16), inference(subsumption_resolution, [], [f2673, f293])).
fof(f2673, plain, ((k6_lattices(sK1) = sK2) | ~ v13_lattices(k1_lattice2(sK1)) | (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ spl29_16), inference(forward_demodulation, [], [f2672, f780])).
fof(f2672, plain, (~ v13_lattices(k1_lattice2(sK1)) | (sK2 = k5_lattices(k1_lattice2(sK1))) | (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ spl29_16), inference(subsumption_resolution, [], [f2671, f288])).
fof(f2671, plain, (~ v13_lattices(k1_lattice2(sK1)) | (sK2 = k5_lattices(k1_lattice2(sK1))) | (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | v3_struct_0(sK1) | ~ spl29_16), inference(subsumption_resolution, [], [f2670, f289])).
fof(f2670, plain, (~ v13_lattices(k1_lattice2(sK1)) | (sK2 = k5_lattices(k1_lattice2(sK1))) | (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ spl29_16), inference(subsumption_resolution, [], [f2669, f290])).
fof(f2669, plain, (~ v13_lattices(k1_lattice2(sK1)) | (sK2 = k5_lattices(k1_lattice2(sK1))) | (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | ~ spl29_16), inference(subsumption_resolution, [], [f2644, f879])).
fof(f2644, plain, (~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | (sK2 = k5_lattices(k1_lattice2(sK1))) | (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f1006, f873])).
fof(f873, plain, m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))), inference(subsumption_resolution, [], [f872, f288])).
fof(f872, plain, (m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f871, f289])).
fof(f871, plain, (m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f870, f290])).
fof(f870, plain, (m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f867, f291])).
fof(f867, plain, (m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | ~ m1_subset_1(sK2, u1_struct_0(sK1)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(superposition, [], [f325, f809])).
fof(f809, plain, (sK2 = k5_filter_2(sK1, sK2)), inference(subsumption_resolution, [], [f808, f288])).
fof(f808, plain, ((sK2 = k5_filter_2(sK1, sK2)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f807, f289])).
fof(f807, plain, ((sK2 = k5_filter_2(sK1, sK2)) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f800, f290])).
fof(f800, plain, ((sK2 = k5_filter_2(sK1, sK2)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f316, f291])).
fof(f316, plain, ! [X0, X1] : (~ m1_subset_1(X1, u1_struct_0(X0)) | (k5_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ! [X0] : (! [X1] : ((k5_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f128])).
fof(f128, plain, ! [X0] : (! [X1] : ((k5_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => (k5_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', d4_filter_2)).
fof(f325, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f141])).
fof(f141, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | (~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ! [X0, X1] : ((m1_subset_1(X1, u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', dt_k5_filter_2)).
fof(f1006, plain, ! [X4, X5] : (~ m1_subset_1(X5, u1_struct_0(k1_lattice2(X4))) | ~ v13_lattices(k1_lattice2(X4)) | ~ v10_lattices(k1_lattice2(X4)) | (k5_lattices(k1_lattice2(X4)) = X5) | (k8_filter_2(X4, sK28(k1_lattice2(X4), X5)) = k14_filter_2(X4, sK28(k1_lattice2(X4), X5))) | ~ l3_lattices(X4) | ~ v10_lattices(X4) | v3_struct_0(X4)), inference(subsumption_resolution, [], [f1005, f360])).
fof(f1005, plain, ! [X4, X5] : ((k5_lattices(k1_lattice2(X4)) = X5) | ~ v13_lattices(k1_lattice2(X4)) | ~ v10_lattices(k1_lattice2(X4)) | v3_struct_0(k1_lattice2(X4)) | ~ m1_subset_1(X5, u1_struct_0(k1_lattice2(X4))) | (k8_filter_2(X4, sK28(k1_lattice2(X4), X5)) = k14_filter_2(X4, sK28(k1_lattice2(X4), X5))) | ~ l3_lattices(X4) | ~ v10_lattices(X4) | v3_struct_0(X4)), inference(subsumption_resolution, [], [f1002, f324])).
fof(f1002, plain, ! [X4, X5] : ((k5_lattices(k1_lattice2(X4)) = X5) | ~ l3_lattices(k1_lattice2(X4)) | ~ v13_lattices(k1_lattice2(X4)) | ~ v10_lattices(k1_lattice2(X4)) | v3_struct_0(k1_lattice2(X4)) | ~ m1_subset_1(X5, u1_struct_0(k1_lattice2(X4))) | (k8_filter_2(X4, sK28(k1_lattice2(X4), X5)) = k14_filter_2(X4, sK28(k1_lattice2(X4), X5))) | ~ l3_lattices(X4) | ~ v10_lattices(X4) | v3_struct_0(X4)), inference(resolution, [], [f464, f441])).
fof(f441, plain, ! [X0, X1] : (~ m1_filter_0(X1, k1_lattice2(X0)) | (k8_filter_2(X0, X1) = k14_filter_2(X0, X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ! [X0, X1] : ((k8_filter_2(X0, X1) = k14_filter_2(X0, X1)) | ~ m1_filter_0(X1, k1_lattice2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f202])).
fof(f202, plain, ! [X0, X1] : ((k8_filter_2(X0, X1) = k14_filter_2(X0, X1)) | (~ m1_filter_0(X1, k1_lattice2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f81])).
fof(f81, plain, ! [X0, X1] : ((m1_filter_0(X1, k1_lattice2(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k8_filter_2(X0, X1) = k14_filter_2(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', redefinition_k14_filter_2)).
fof(f1268, plain, spl29_60, inference(avatar_contradiction_clause, [], [f1267])).
fof(f1267, plain, ($false | spl29_60), inference(subsumption_resolution, [], [f1266, f473])).
fof(f1266, plain, (~ l1_lattices(sK1) | spl29_60), inference(resolution, [], [f1233, f342])).
fof(f342, plain, ! [X0] : (v1_funct_1(u1_lattices(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f163])).
fof(f1233, plain, (~ v1_funct_1(u1_lattices(sK1)) | spl29_60), inference(avatar_component_clause, [], [f1231])).
fof(f1231, plain, (spl29_60 <=> v1_funct_1(u1_lattices(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_60])])).
fof(f1265, plain, (~ spl29_60 | ~ spl29_61 | ~ spl29_62 | ~ spl29_63 | ~ spl29_64 | ~ spl29_65 | spl29_68), inference(avatar_split_clause, [], [f1229, f1263, f1251, f1247, f1243, f1239, f1235, f1231])).
fof(f1229, plain, ! [X17, X15, X16] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15) | ~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u2_lattices(sK1)) | ~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u1_lattices(sK1))), inference(superposition, [], [f391, f710])).
fof(f710, plain, (k1_lattice2(sK1) = g3_lattices(u1_struct_0(sK1), u1_lattices(sK1), u2_lattices(sK1))), inference(resolution, [], [f315, f290])).
fof(f315, plain, ! [X0] : (~ l3_lattices(X0) | (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ! [X0] : ((k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ! [X0] : (l3_lattices(X0) => (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', d2_lattice2)).
fof(f391, plain, ! [X4, X2, X0, X5, X3, X1] : (~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) | (X0 = X3) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f196])).
fof(f196, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(flattening, [], [f195])).
fof(f195, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | (~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1))), inference(ennf_transformation, [], [f66])).
fof(f66, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X2) & m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X1)) => ! [X3, X4, X5] : ((g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) => ((X2 = X5) & (X1 = X4) & (X0 = X3)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', free_g3_lattices)).
fof(f926, plain, spl29_17, inference(avatar_contradiction_clause, [], [f925])).
fof(f925, plain, ($false | spl29_17), inference(subsumption_resolution, [], [f924, f290])).
fof(f924, plain, (~ l3_lattices(sK1) | spl29_17), inference(resolution, [], [f884, f324])).
fof(f884, plain, (~ l3_lattices(k1_lattice2(sK1)) | spl29_17), inference(avatar_component_clause, [], [f882])).
fof(f916, plain, spl29_16, inference(avatar_contradiction_clause, [], [f915])).
fof(f915, plain, ($false | spl29_16), inference(subsumption_resolution, [], [f914, f618])).
fof(f618, plain, sP0(sK1), inference(subsumption_resolution, [], [f617, f288])).
fof(f617, plain, (sP0(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f611, f289])).
fof(f611, plain, (sP0(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f390, f290])).
fof(f390, plain, ! [X0] : (~ l3_lattices(X0) | sP0(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f227])).
fof(f227, plain, ! [X0] : (sP0(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(definition_folding, [], [f194, e226])).
fof(f226, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(usedef, [], [e226])).
fof(e226, plain, ! [X0] : (sP0(X0) <=> (v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f194, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f193])).
fof(f193, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f65])).
fof(f65, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT309+1.p', fc6_lattice2)).
fof(f914, plain, (~ sP0(sK1) | spl29_16), inference(resolution, [], [f880, f389])).
fof(f389, plain, ! [X0] : (v10_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f253])).
fof(f253, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(nnf_transformation, [], [f226])).
fof(f880, plain, (~ v10_lattices(k1_lattice2(sK1)) | spl29_16), inference(avatar_component_clause, [], [f878])).
fof(f823, plain, ~ spl29_13, inference(avatar_contradiction_clause, [], [f822])).
fof(f822, plain, ($false | ~ spl29_13), inference(subsumption_resolution, [], [f818, f618])).
fof(f818, plain, (~ sP0(sK1) | ~ spl29_13), inference(resolution, [], [f788, f381])).
fof(f381, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f253])).
fof(f788, plain, (v3_struct_0(k1_lattice2(sK1)) | ~ spl29_13), inference(avatar_component_clause, [], [f786])).