fof(f493501, plain, $false, inference(avatar_sat_refutation, [], [f976, f981, f1005, f1025, f2088, f15383, f15385, f15389, f15392, f15395, f15472, f15501, f20611, f20616, f493500])).
fof(f493500, plain, (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_78 | ~ spl33_91 | ~ spl33_92), inference(avatar_contradiction_clause, [], [f493499])).
fof(f493499, plain, ($false | (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_78 | ~ spl33_91 | ~ spl33_92)), inference(subsumption_resolution, [], [f493498, f341])).
fof(f341, plain, ~ (sK2 = k3_lattices(sK1, sK2, sK3)), inference(cnf_transformation, [], [f268])).
fof(f268, plain, (((u1_struct_0(sK1) = k18_filter_2(sK1, sK2)) & (~ (sK2 = k3_lattices(sK1, sK2, sK3)) & m1_subset_1(sK3, u1_struct_0(sK1))) & m1_subset_1(sK2, u1_struct_0(sK1))) & l3_lattices(sK1) & v10_lattices(sK1) & ~ v3_struct_0(sK1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2, sK3])], [f132, f267, f266, f265])).
fof(f265, plain, (? [X0] : (? [X1] : ((u1_struct_0(X0) = k18_filter_2(X0, X1)) & ? [X2] : (~ (k3_lattices(X0, X1, X2) = X1) & m1_subset_1(X2, u1_struct_0(X0))) & m1_subset_1(X1, u1_struct_0(X0))) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (? [X1] : ((u1_struct_0(sK1) = k18_filter_2(sK1, X1)) & ? [X2] : (~ (k3_lattices(sK1, X1, X2) = X1) & m1_subset_1(X2, u1_struct_0(sK1))) & m1_subset_1(X1, u1_struct_0(sK1))) & l3_lattices(sK1) & v10_lattices(sK1) & ~ v3_struct_0(sK1))), introduced(choice_axiom, [])).
fof(f266, plain, (? [X1] : ((u1_struct_0(sK1) = k18_filter_2(sK1, X1)) & ? [X2] : (~ (k3_lattices(sK1, X1, X2) = X1) & m1_subset_1(X2, u1_struct_0(sK1))) & m1_subset_1(X1, u1_struct_0(sK1))) => ((u1_struct_0(sK1) = k18_filter_2(sK1, sK2)) & ? [X2] : (~ (sK2 = k3_lattices(sK1, sK2, X2)) & m1_subset_1(X2, u1_struct_0(sK1))) & m1_subset_1(sK2, u1_struct_0(sK1)))), introduced(choice_axiom, [])).
fof(f267, plain, (? [X2] : (~ (sK2 = k3_lattices(sK1, sK2, X2)) & m1_subset_1(X2, u1_struct_0(sK1))) => (~ (sK2 = k3_lattices(sK1, sK2, sK3)) & m1_subset_1(sK3, u1_struct_0(sK1)))), introduced(choice_axiom, [])).
fof(f132, plain, ? [X0] : (? [X1] : ((u1_struct_0(X0) = k18_filter_2(X0, X1)) & ? [X2] : (~ (k3_lattices(X0, X1, X2) = X1) & m1_subset_1(X2, u1_struct_0(X0))) & m1_subset_1(X1, u1_struct_0(X0))) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f131])).
fof(f131, plain, ? [X0] : (? [X1] : (((u1_struct_0(X0) = k18_filter_2(X0, X1)) & ? [X2] : (~ (k3_lattices(X0, X1, X2) = X1) & m1_subset_1(X2, u1_struct_0(X0)))) & m1_subset_1(X1, u1_struct_0(X0))) & (l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ~ ((u1_struct_0(X0) = k18_filter_2(X0, X1)) & ~ ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (k3_lattices(X0, X1, X2) = X1))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ~ ((u1_struct_0(X0) = k18_filter_2(X0, X1)) & ~ ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (k3_lattices(X0, X1, X2) = X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', t35_filter_2)).
fof(f493498, plain, ((sK2 = k3_lattices(sK1, sK2, sK3)) | (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_78 | ~ spl33_91 | ~ spl33_92)), inference(backward_demodulation, [], [f3316, f493497])).
fof(f493497, plain, ((sK2 = k2_binop_1(u1_struct_0(sK1), u1_struct_0(sK1), u1_struct_0(sK1), u2_lattices(sK1), sK3, sK2)) | (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_78 | ~ spl33_91 | ~ spl33_92)), inference(forward_demodulation, [], [f493355, f470246])).
fof(f470246, plain, ((sK2 = k2_lattices(k1_lattice2(sK1), sK3, sK2)) | (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_91 | ~ spl33_92)), inference(forward_demodulation, [], [f470177, f440267])).
fof(f440267, plain, ((sK2 = k4_lattices(k1_lattice2(sK1), sK3, sK2)) | (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_91 | ~ spl33_92)), inference(forward_demodulation, [], [f433462, f366885])).
fof(f366885, plain, ((sK2 = k4_lattices(k1_lattice2(sK1), sK2, sK3)) | (~ spl33_91 | ~ spl33_92)), inference(resolution, [], [f21010, f340])).
fof(f340, plain, m1_subset_1(sK3, u1_struct_0(sK1)), inference(cnf_transformation, [], [f268])).
fof(f21010, plain, (! [X0] : (~ m1_subset_1(X0, u1_struct_0(sK1)) | (sK2 = k4_lattices(k1_lattice2(sK1), sK2, X0))) | (~ spl33_91 | ~ spl33_92)), inference(forward_demodulation, [], [f2779, f2782])).
fof(f2782, plain, ((u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | ~ spl33_92), inference(avatar_component_clause, [], [f2781])).
fof(f2781, plain, (spl33_92 <=> (u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1)))), introduced(avatar_definition, [new_symbols(naming, [spl33_92])])).
fof(f2779, plain, (! [X0] : ((sK2 = k4_lattices(k1_lattice2(sK1), sK2, X0)) | ~ m1_subset_1(X0, u1_struct_0(k1_lattice2(sK1)))) | ~ spl33_91), inference(avatar_component_clause, [], [f2778])).
fof(f2778, plain, (spl33_91 <=> ! [X0] : ((sK2 = k4_lattices(k1_lattice2(sK1), sK2, X0)) | ~ m1_subset_1(X0, u1_struct_0(k1_lattice2(sK1))))), introduced(avatar_definition, [new_symbols(naming, [spl33_91])])).
fof(f433462, plain, ((k4_lattices(k1_lattice2(sK1), sK2, sK3) = k4_lattices(k1_lattice2(sK1), sK3, sK2)) | (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_92)), inference(resolution, [], [f20628, f340])).
fof(f20628, plain, (! [X27] : (~ m1_subset_1(X27, u1_struct_0(sK1)) | (k4_lattices(k1_lattice2(sK1), X27, sK2) = k4_lattices(k1_lattice2(sK1), sK2, X27))) | (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_92)), inference(backward_demodulation, [], [f1347, f2782])).
fof(f1347, plain, (! [X27] : ((k4_lattices(k1_lattice2(sK1), X27, sK2) = k4_lattices(k1_lattice2(sK1), sK2, X27)) | ~ m1_subset_1(X27, u1_struct_0(k1_lattice2(sK1)))) | (spl33_14 | ~ spl33_15 | ~ spl33_16)), inference(subsumption_resolution, [], [f1346, f927])).
fof(f927, plain, (~ v3_struct_0(k1_lattice2(sK1)) | spl33_14), inference(avatar_component_clause, [], [f926])).
fof(f926, plain, (spl33_14 <=> v3_struct_0(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl33_14])])).
fof(f1346, plain, (! [X27] : ((k4_lattices(k1_lattice2(sK1), X27, sK2) = k4_lattices(k1_lattice2(sK1), sK2, X27)) | ~ m1_subset_1(X27, u1_struct_0(k1_lattice2(sK1))) | v3_struct_0(k1_lattice2(sK1))) | (spl33_14 | ~ spl33_15 | ~ spl33_16)), inference(subsumption_resolution, [], [f1345, f1082])).
fof(f1082, plain, (v6_lattices(k1_lattice2(sK1)) | (spl33_14 | ~ spl33_15 | ~ spl33_16)), inference(subsumption_resolution, [], [f1081, f927])).
fof(f1081, plain, (v3_struct_0(k1_lattice2(sK1)) | v6_lattices(k1_lattice2(sK1)) | (~ spl33_15 | ~ spl33_16)), inference(subsumption_resolution, [], [f1069, f931])).
fof(f931, plain, (v10_lattices(k1_lattice2(sK1)) | ~ spl33_15), inference(avatar_component_clause, [], [f930])).
fof(f930, plain, (spl33_15 <=> v10_lattices(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl33_15])])).
fof(f1069, plain, (~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | v6_lattices(k1_lattice2(sK1)) | ~ spl33_16), inference(resolution, [], [f935, f349])).
fof(f349, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v6_lattices(X0)), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ! [X0] : ((v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f137])).
fof(f137, plain, ! [X0] : (((v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)) | (~ v10_lattices(X0) | v3_struct_0(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ! [X0] : (l3_lattices(X0) => ((v10_lattices(X0) & ~ v3_struct_0(X0)) => (v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', cc1_lattices)).
fof(f935, plain, (l3_lattices(k1_lattice2(sK1)) | ~ spl33_16), inference(avatar_component_clause, [], [f934])).
fof(f934, plain, (spl33_16 <=> l3_lattices(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl33_16])])).
fof(f1345, plain, (! [X27] : ((k4_lattices(k1_lattice2(sK1), X27, sK2) = k4_lattices(k1_lattice2(sK1), sK2, X27)) | ~ m1_subset_1(X27, u1_struct_0(k1_lattice2(sK1))) | ~ v6_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))) | ~ spl33_16), inference(subsumption_resolution, [], [f1328, f1074])).
fof(f1074, plain, (l1_lattices(k1_lattice2(sK1)) | ~ spl33_16), inference(resolution, [], [f935, f385])).
fof(f385, plain, ! [X0] : (~ l3_lattices(X0) | l1_lattices(X0)), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ! [X0] : ((l2_lattices(X0) & l1_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f44])).
fof(f44, plain, ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', dt_l3_lattices)).
fof(f1328, plain, ! [X27] : ((k4_lattices(k1_lattice2(sK1), X27, sK2) = k4_lattices(k1_lattice2(sK1), sK2, X27)) | ~ m1_subset_1(X27, u1_struct_0(k1_lattice2(sK1))) | ~ l1_lattices(k1_lattice2(sK1)) | ~ v6_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(resolution, [], [f361, f916])).
fof(f916, plain, m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))), inference(subsumption_resolution, [], [f915, f336])).
fof(f336, plain, ~ v3_struct_0(sK1), inference(cnf_transformation, [], [f268])).
fof(f915, plain, (m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f914, f337])).
fof(f337, plain, v10_lattices(sK1), inference(cnf_transformation, [], [f268])).
fof(f914, plain, (m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f913, f338])).
fof(f338, plain, l3_lattices(sK1), inference(cnf_transformation, [], [f268])).
fof(f913, plain, (m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f905, f339])).
fof(f339, plain, m1_subset_1(sK2, u1_struct_0(sK1)), inference(cnf_transformation, [], [f268])).
fof(f905, plain, (m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | ~ m1_subset_1(sK2, u1_struct_0(sK1)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(superposition, [], [f382, f854])).
fof(f854, plain, (sK2 = k5_filter_2(sK1, sK2)), inference(subsumption_resolution, [], [f853, f336])).
fof(f853, plain, ((sK2 = k5_filter_2(sK1, sK2)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f852, f337])).
fof(f852, plain, ((sK2 = k5_filter_2(sK1, sK2)) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f849, f338])).
fof(f849, plain, ((sK2 = k5_filter_2(sK1, sK2)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f367, f339])).
fof(f367, plain, ! [X0, X1] : (~ m1_subset_1(X1, u1_struct_0(X0)) | (k5_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ! [X0] : (! [X1] : ((k5_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f158])).
fof(f158, plain, ! [X0] : (! [X1] : ((k5_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => (k5_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', d4_filter_2)).
fof(f382, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f181])).
fof(f181, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | (~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f40])).
fof(f40, plain, ! [X0, X1] : ((m1_subset_1(X1, u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', dt_k5_filter_2)).
fof(f361, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, u1_struct_0(X0)) | (k4_lattices(X0, X1, X2) = k4_lattices(X0, X2, X1)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l1_lattices(X0) | ~ v6_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ! [X0, X1, X2] : ((k4_lattices(X0, X1, X2) = k4_lattices(X0, X2, X1)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l1_lattices(X0) | ~ v6_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f147])).
fof(f147, plain, ! [X0, X1, X2] : ((k4_lattices(X0, X1, X2) = k4_lattices(X0, X2, X1)) | (~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l1_lattices(X0) | ~ v6_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0, X1, X2] : ((m1_subset_1(X2, u1_struct_0(X0)) & m1_subset_1(X1, u1_struct_0(X0)) & l1_lattices(X0) & v6_lattices(X0) & ~ v3_struct_0(X0)) => (k4_lattices(X0, X1, X2) = k4_lattices(X0, X2, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', commutativity_k4_lattices)).
fof(f470177, plain, ((k4_lattices(k1_lattice2(sK1), sK3, sK2) = k2_lattices(k1_lattice2(sK1), sK3, sK2)) | (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_92)), inference(resolution, [], [f20631, f340])).
fof(f20631, plain, (! [X3] : (~ m1_subset_1(X3, u1_struct_0(sK1)) | (k4_lattices(k1_lattice2(sK1), X3, sK2) = k2_lattices(k1_lattice2(sK1), X3, sK2))) | (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_92)), inference(backward_demodulation, [], [f1467, f2782])).
fof(f1467, plain, (! [X3] : ((k4_lattices(k1_lattice2(sK1), X3, sK2) = k2_lattices(k1_lattice2(sK1), X3, sK2)) | ~ m1_subset_1(X3, u1_struct_0(k1_lattice2(sK1)))) | (spl33_14 | ~ spl33_15 | ~ spl33_16)), inference(subsumption_resolution, [], [f1466, f927])).
fof(f1466, plain, (! [X3] : ((k4_lattices(k1_lattice2(sK1), X3, sK2) = k2_lattices(k1_lattice2(sK1), X3, sK2)) | ~ m1_subset_1(X3, u1_struct_0(k1_lattice2(sK1))) | v3_struct_0(k1_lattice2(sK1))) | (spl33_14 | ~ spl33_15 | ~ spl33_16)), inference(subsumption_resolution, [], [f1465, f1082])).
fof(f1465, plain, (! [X3] : ((k4_lattices(k1_lattice2(sK1), X3, sK2) = k2_lattices(k1_lattice2(sK1), X3, sK2)) | ~ m1_subset_1(X3, u1_struct_0(k1_lattice2(sK1))) | ~ v6_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))) | ~ spl33_16), inference(subsumption_resolution, [], [f1445, f1074])).
fof(f1445, plain, ! [X3] : ((k4_lattices(k1_lattice2(sK1), X3, sK2) = k2_lattices(k1_lattice2(sK1), X3, sK2)) | ~ m1_subset_1(X3, u1_struct_0(k1_lattice2(sK1))) | ~ l1_lattices(k1_lattice2(sK1)) | ~ v6_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(resolution, [], [f916, f498])).
fof(f498, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, u1_struct_0(X0)) | (k4_lattices(X0, X1, X2) = k2_lattices(X0, X1, X2)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l1_lattices(X0) | ~ v6_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f241])).
fof(f241, plain, ! [X0, X1, X2] : ((k4_lattices(X0, X1, X2) = k2_lattices(X0, X1, X2)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l1_lattices(X0) | ~ v6_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f240])).
fof(f240, plain, ! [X0, X1, X2] : ((k4_lattices(X0, X1, X2) = k2_lattices(X0, X1, X2)) | (~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l1_lattices(X0) | ~ v6_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f101])).
fof(f101, plain, ! [X0, X1, X2] : ((m1_subset_1(X2, u1_struct_0(X0)) & m1_subset_1(X1, u1_struct_0(X0)) & l1_lattices(X0) & v6_lattices(X0) & ~ v3_struct_0(X0)) => (k4_lattices(X0, X1, X2) = k2_lattices(X0, X1, X2))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', redefinition_k4_lattices)).
fof(f493355, plain, ((k2_binop_1(u1_struct_0(sK1), u1_struct_0(sK1), u1_struct_0(sK1), u2_lattices(sK1), sK3, sK2) = k2_lattices(k1_lattice2(sK1), sK3, sK2)) | (spl33_14 | ~ spl33_16 | ~ spl33_78 | ~ spl33_92)), inference(resolution, [], [f20936, f340])).
fof(f20936, plain, (! [X27] : (~ m1_subset_1(X27, u1_struct_0(sK1)) | (k2_lattices(k1_lattice2(sK1), X27, sK2) = k2_binop_1(u1_struct_0(sK1), u1_struct_0(sK1), u1_struct_0(sK1), u2_lattices(sK1), X27, sK2))) | (spl33_14 | ~ spl33_16 | ~ spl33_78 | ~ spl33_92)), inference(forward_demodulation, [], [f20831, f2782])).
fof(f20831, plain, (! [X27] : (~ m1_subset_1(X27, u1_struct_0(sK1)) | (k2_lattices(k1_lattice2(sK1), X27, sK2) = k2_binop_1(u1_struct_0(k1_lattice2(sK1)), u1_struct_0(k1_lattice2(sK1)), u1_struct_0(k1_lattice2(sK1)), u2_lattices(sK1), X27, sK2))) | (spl33_14 | ~ spl33_16 | ~ spl33_78 | ~ spl33_92)), inference(backward_demodulation, [], [f20554, f2782])).
fof(f20554, plain, (! [X27] : ((k2_lattices(k1_lattice2(sK1), X27, sK2) = k2_binop_1(u1_struct_0(k1_lattice2(sK1)), u1_struct_0(k1_lattice2(sK1)), u1_struct_0(k1_lattice2(sK1)), u2_lattices(sK1), X27, sK2)) | ~ m1_subset_1(X27, u1_struct_0(k1_lattice2(sK1)))) | (spl33_14 | ~ spl33_16 | ~ spl33_78)), inference(backward_demodulation, [], [f1618, f20516])).
fof(f20516, plain, ((u2_lattices(sK1) = u1_lattices(k1_lattice2(sK1))) | ~ spl33_78), inference(trivial_inequality_removal, [], [f20499])).
fof(f20499, plain, (~ (k1_lattice2(sK1) = k1_lattice2(sK1)) | (u2_lattices(sK1) = u1_lattices(k1_lattice2(sK1))) | ~ spl33_78), inference(superposition, [], [f2076, f3002])).
fof(f3002, plain, (k1_lattice2(sK1) = g3_lattices(u1_struct_0(k1_lattice2(sK1)), u2_lattices(k1_lattice2(sK1)), u1_lattices(k1_lattice2(sK1)))), inference(resolution, [], [f799, f338])).
fof(f799, plain, ! [X1] : (~ l3_lattices(X1) | (k1_lattice2(X1) = g3_lattices(u1_struct_0(k1_lattice2(X1)), u2_lattices(k1_lattice2(X1)), u1_lattices(k1_lattice2(X1))))), inference(subsumption_resolution, [], [f795, f374])).
fof(f374, plain, ! [X0] : (l3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ! [X0] : ((l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f26])).
fof(f26, plain, ! [X0] : (l3_lattices(X0) => (l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', dt_k1_lattice2)).
fof(f795, plain, ! [X1] : ((k1_lattice2(X1) = g3_lattices(u1_struct_0(k1_lattice2(X1)), u2_lattices(k1_lattice2(X1)), u1_lattices(k1_lattice2(X1)))) | ~ l3_lattices(k1_lattice2(X1)) | ~ l3_lattices(X1)), inference(resolution, [], [f343, f373])).
fof(f373, plain, ! [X0] : (v3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f166])).
fof(f343, plain, ! [X0] : (~ v3_lattices(X0) | (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ! [X0] : ((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f133])).
fof(f133, plain, ! [X0] : (((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (l3_lattices(X0) => (v3_lattices(X0) => (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', abstractness_v3_lattices)).
fof(f2076, plain, (! [X4, X5, X3] : (~ (g3_lattices(X3, X4, X5) = k1_lattice2(sK1)) | (u2_lattices(sK1) = X5)) | ~ spl33_78), inference(avatar_component_clause, [], [f2075])).
fof(f2075, plain, (spl33_78 <=> ! [X3, X5, X4] : (~ (g3_lattices(X3, X4, X5) = k1_lattice2(sK1)) | (u2_lattices(sK1) = X5))), introduced(avatar_definition, [new_symbols(naming, [spl33_78])])).
fof(f1618, plain, (! [X27] : ((k2_lattices(k1_lattice2(sK1), X27, sK2) = k2_binop_1(u1_struct_0(k1_lattice2(sK1)), u1_struct_0(k1_lattice2(sK1)), u1_struct_0(k1_lattice2(sK1)), u1_lattices(k1_lattice2(sK1)), X27, sK2)) | ~ m1_subset_1(X27, u1_struct_0(k1_lattice2(sK1)))) | (spl33_14 | ~ spl33_16)), inference(subsumption_resolution, [], [f1617, f927])).
fof(f1617, plain, (! [X27] : ((k2_lattices(k1_lattice2(sK1), X27, sK2) = k2_binop_1(u1_struct_0(k1_lattice2(sK1)), u1_struct_0(k1_lattice2(sK1)), u1_struct_0(k1_lattice2(sK1)), u1_lattices(k1_lattice2(sK1)), X27, sK2)) | ~ m1_subset_1(X27, u1_struct_0(k1_lattice2(sK1))) | v3_struct_0(k1_lattice2(sK1))) | ~ spl33_16), inference(subsumption_resolution, [], [f1602, f1074])).
fof(f1602, plain, ! [X27] : ((k2_lattices(k1_lattice2(sK1), X27, sK2) = k2_binop_1(u1_struct_0(k1_lattice2(sK1)), u1_struct_0(k1_lattice2(sK1)), u1_struct_0(k1_lattice2(sK1)), u1_lattices(k1_lattice2(sK1)), X27, sK2)) | ~ m1_subset_1(X27, u1_struct_0(k1_lattice2(sK1))) | ~ l1_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(resolution, [], [f365, f916])).
fof(f365, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, u1_struct_0(X0)) | (k2_lattices(X0, X1, X2) = k2_binop_1(u1_struct_0(X0), u1_struct_0(X0), u1_struct_0(X0), u1_lattices(X0), X1, X2)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l1_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ! [X0] : (! [X1] : (! [X2] : ((k2_lattices(X0, X1, X2) = k2_binop_1(u1_struct_0(X0), u1_struct_0(X0), u1_struct_0(X0), u1_lattices(X0), X1, X2)) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l1_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f154])).
fof(f154, plain, ! [X0] : (! [X1] : (! [X2] : ((k2_lattices(X0, X1, X2) = k2_binop_1(u1_struct_0(X0), u1_struct_0(X0), u1_struct_0(X0), u1_lattices(X0), X1, X2)) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l1_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0] : ((l1_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (k2_lattices(X0, X1, X2) = k2_binop_1(u1_struct_0(X0), u1_struct_0(X0), u1_struct_0(X0), u1_lattices(X0), X1, X2))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', d2_lattices)).
fof(f3316, plain, (k3_lattices(sK1, sK2, sK3) = k2_binop_1(u1_struct_0(sK1), u1_struct_0(sK1), u1_struct_0(sK1), u2_lattices(sK1), sK3, sK2)), inference(forward_demodulation, [], [f3303, f2590])).
fof(f2590, plain, (k3_lattices(sK1, sK2, sK3) = k1_lattices(sK1, sK3, sK2)), inference(forward_demodulation, [], [f2581, f2399])).
fof(f2399, plain, (k3_lattices(sK1, sK2, sK3) = k3_lattices(sK1, sK3, sK2)), inference(resolution, [], [f1275, f340])).
fof(f1275, plain, ! [X28] : (~ m1_subset_1(X28, u1_struct_0(sK1)) | (k3_lattices(sK1, X28, sK2) = k3_lattices(sK1, sK2, X28))), inference(subsumption_resolution, [], [f1274, f336])).
fof(f1274, plain, ! [X28] : ((k3_lattices(sK1, X28, sK2) = k3_lattices(sK1, sK2, X28)) | ~ m1_subset_1(X28, u1_struct_0(sK1)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1273, f583])).
fof(f583, plain, v4_lattices(sK1), inference(subsumption_resolution, [], [f582, f336])).
fof(f582, plain, (v3_struct_0(sK1) | v4_lattices(sK1)), inference(subsumption_resolution, [], [f577, f337])).
fof(f577, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | v4_lattices(sK1)), inference(resolution, [], [f347, f338])).
fof(f347, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v4_lattices(X0)), inference(cnf_transformation, [], [f138])).
fof(f1273, plain, ! [X28] : ((k3_lattices(sK1, X28, sK2) = k3_lattices(sK1, sK2, X28)) | ~ m1_subset_1(X28, u1_struct_0(sK1)) | ~ v4_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1255, f540])).
fof(f540, plain, l2_lattices(sK1), inference(resolution, [], [f386, f338])).
fof(f386, plain, ! [X0] : (~ l3_lattices(X0) | l2_lattices(X0)), inference(cnf_transformation, [], [f185])).
fof(f1255, plain, ! [X28] : ((k3_lattices(sK1, X28, sK2) = k3_lattices(sK1, sK2, X28)) | ~ m1_subset_1(X28, u1_struct_0(sK1)) | ~ l2_lattices(sK1) | ~ v4_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f360, f339])).
fof(f360, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, u1_struct_0(X0)) | (k3_lattices(X0, X1, X2) = k3_lattices(X0, X2, X1)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l2_lattices(X0) | ~ v4_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ! [X0, X1, X2] : ((k3_lattices(X0, X1, X2) = k3_lattices(X0, X2, X1)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l2_lattices(X0) | ~ v4_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f145])).
fof(f145, plain, ! [X0, X1, X2] : ((k3_lattices(X0, X1, X2) = k3_lattices(X0, X2, X1)) | (~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l2_lattices(X0) | ~ v4_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ! [X0, X1, X2] : ((m1_subset_1(X2, u1_struct_0(X0)) & m1_subset_1(X1, u1_struct_0(X0)) & l2_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)) => (k3_lattices(X0, X1, X2) = k3_lattices(X0, X2, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', commutativity_k3_lattices)).
fof(f2581, plain, (k3_lattices(sK1, sK3, sK2) = k1_lattices(sK1, sK3, sK2)), inference(resolution, [], [f1406, f340])).
fof(f1406, plain, ! [X27] : (~ m1_subset_1(X27, u1_struct_0(sK1)) | (k3_lattices(sK1, X27, sK2) = k1_lattices(sK1, X27, sK2))), inference(subsumption_resolution, [], [f1405, f336])).
fof(f1405, plain, ! [X27] : ((k3_lattices(sK1, X27, sK2) = k1_lattices(sK1, X27, sK2)) | ~ m1_subset_1(X27, u1_struct_0(sK1)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1404, f583])).
fof(f1404, plain, ! [X27] : ((k3_lattices(sK1, X27, sK2) = k1_lattices(sK1, X27, sK2)) | ~ m1_subset_1(X27, u1_struct_0(sK1)) | ~ v4_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1389, f540])).
fof(f1389, plain, ! [X27] : ((k3_lattices(sK1, X27, sK2) = k1_lattices(sK1, X27, sK2)) | ~ m1_subset_1(X27, u1_struct_0(sK1)) | ~ l2_lattices(sK1) | ~ v4_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f497, f339])).
fof(f497, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, u1_struct_0(X0)) | (k3_lattices(X0, X1, X2) = k1_lattices(X0, X1, X2)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l2_lattices(X0) | ~ v4_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f239])).
fof(f239, plain, ! [X0, X1, X2] : ((k3_lattices(X0, X1, X2) = k1_lattices(X0, X1, X2)) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l2_lattices(X0) | ~ v4_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f238])).
fof(f238, plain, ! [X0, X1, X2] : ((k3_lattices(X0, X1, X2) = k1_lattices(X0, X1, X2)) | (~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l2_lattices(X0) | ~ v4_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f100])).
fof(f100, plain, ! [X0, X1, X2] : ((m1_subset_1(X2, u1_struct_0(X0)) & m1_subset_1(X1, u1_struct_0(X0)) & l2_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)) => (k3_lattices(X0, X1, X2) = k1_lattices(X0, X1, X2))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', redefinition_k3_lattices)).
fof(f3303, plain, (k1_lattices(sK1, sK3, sK2) = k2_binop_1(u1_struct_0(sK1), u1_struct_0(sK1), u1_struct_0(sK1), u2_lattices(sK1), sK3, sK2)), inference(resolution, [], [f1587, f340])).
fof(f1587, plain, ! [X29] : (~ m1_subset_1(X29, u1_struct_0(sK1)) | (k1_lattices(sK1, X29, sK2) = k2_binop_1(u1_struct_0(sK1), u1_struct_0(sK1), u1_struct_0(sK1), u2_lattices(sK1), X29, sK2))), inference(subsumption_resolution, [], [f1586, f336])).
fof(f1586, plain, ! [X29] : ((k1_lattices(sK1, X29, sK2) = k2_binop_1(u1_struct_0(sK1), u1_struct_0(sK1), u1_struct_0(sK1), u2_lattices(sK1), X29, sK2)) | ~ m1_subset_1(X29, u1_struct_0(sK1)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1569, f540])).
fof(f1569, plain, ! [X29] : ((k1_lattices(sK1, X29, sK2) = k2_binop_1(u1_struct_0(sK1), u1_struct_0(sK1), u1_struct_0(sK1), u2_lattices(sK1), X29, sK2)) | ~ m1_subset_1(X29, u1_struct_0(sK1)) | ~ l2_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f363, f339])).
fof(f363, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, u1_struct_0(X0)) | (k1_lattices(X0, X1, X2) = k2_binop_1(u1_struct_0(X0), u1_struct_0(X0), u1_struct_0(X0), u2_lattices(X0), X1, X2)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l2_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ! [X0] : (! [X1] : (! [X2] : ((k1_lattices(X0, X1, X2) = k2_binop_1(u1_struct_0(X0), u1_struct_0(X0), u1_struct_0(X0), u2_lattices(X0), X1, X2)) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l2_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f151])).
fof(f151, plain, ! [X0] : (! [X1] : (! [X2] : ((k1_lattices(X0, X1, X2) = k2_binop_1(u1_struct_0(X0), u1_struct_0(X0), u1_struct_0(X0), u2_lattices(X0), X1, X2)) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l2_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ! [X0] : ((l2_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (k1_lattices(X0, X1, X2) = k2_binop_1(u1_struct_0(X0), u1_struct_0(X0), u1_struct_0(X0), u2_lattices(X0), X1, X2))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', d1_lattices)).
fof(f20616, plain, (spl33_91 | ~ spl33_92 | spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_22), inference(avatar_split_clause, [], [f20615, f978, f934, f930, f926, f2781, f2778])).
fof(f978, plain, (spl33_22 <=> (k2_filter_0(k1_lattice2(sK1), sK2) = k2_filter_2(k1_lattice2(sK1), sK2))), introduced(avatar_definition, [new_symbols(naming, [spl33_22])])).
fof(f20615, plain, (! [X0] : (~ (u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | (sK2 = k4_lattices(k1_lattice2(sK1), sK2, X0)) | ~ m1_subset_1(X0, u1_struct_0(k1_lattice2(sK1)))) | (spl33_14 | ~ spl33_15 | ~ spl33_16 | ~ spl33_22)), inference(subsumption_resolution, [], [f20614, f927])).
fof(f20614, plain, (! [X0] : (~ (u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | (sK2 = k4_lattices(k1_lattice2(sK1), sK2, X0)) | ~ m1_subset_1(X0, u1_struct_0(k1_lattice2(sK1))) | v3_struct_0(k1_lattice2(sK1))) | (~ spl33_15 | ~ spl33_16 | ~ spl33_22)), inference(subsumption_resolution, [], [f20613, f931])).
fof(f20613, plain, (! [X0] : (~ (u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | (sK2 = k4_lattices(k1_lattice2(sK1), sK2, X0)) | ~ m1_subset_1(X0, u1_struct_0(k1_lattice2(sK1))) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))) | (~ spl33_16 | ~ spl33_22)), inference(subsumption_resolution, [], [f20612, f935])).
fof(f20612, plain, (! [X0] : (~ (u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | (sK2 = k4_lattices(k1_lattice2(sK1), sK2, X0)) | ~ m1_subset_1(X0, u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))) | ~ spl33_22), inference(subsumption_resolution, [], [f6267, f916])).
fof(f6267, plain, (! [X0] : (~ (u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | (sK2 = k4_lattices(k1_lattice2(sK1), sK2, X0)) | ~ m1_subset_1(X0, u1_struct_0(k1_lattice2(sK1))) | ~ m1_subset_1(sK2, u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))) | ~ spl33_22), inference(superposition, [], [f508, f1053])).
fof(f1053, plain, ((u1_struct_0(sK1) = k2_filter_0(k1_lattice2(sK1), sK2)) | ~ spl33_22), inference(backward_demodulation, [], [f980, f1052])).
fof(f1052, plain, (u1_struct_0(sK1) = k2_filter_2(k1_lattice2(sK1), sK2)), inference(forward_demodulation, [], [f1051, f342])).
fof(f342, plain, (u1_struct_0(sK1) = k18_filter_2(sK1, sK2)), inference(cnf_transformation, [], [f268])).
fof(f1051, plain, (k18_filter_2(sK1, sK2) = k2_filter_2(k1_lattice2(sK1), sK2)), inference(forward_demodulation, [], [f1050, f854])).
fof(f1050, plain, (k18_filter_2(sK1, sK2) = k2_filter_2(k1_lattice2(sK1), k5_filter_2(sK1, sK2))), inference(subsumption_resolution, [], [f1049, f336])).
fof(f1049, plain, ((k18_filter_2(sK1, sK2) = k2_filter_2(k1_lattice2(sK1), k5_filter_2(sK1, sK2))) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1048, f337])).
fof(f1048, plain, ((k18_filter_2(sK1, sK2) = k2_filter_2(k1_lattice2(sK1), k5_filter_2(sK1, sK2))) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f1039, f338])).
fof(f1039, plain, ((k18_filter_2(sK1, sK2) = k2_filter_2(k1_lattice2(sK1), k5_filter_2(sK1, sK2))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f512, f339])).
fof(f512, plain, ! [X0, X1] : (~ m1_subset_1(X1, u1_struct_0(X0)) | (k18_filter_2(X0, X1) = k2_filter_2(k1_lattice2(X0), k5_filter_2(X0, X1))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f255])).
fof(f255, plain, ! [X0] : (! [X1] : (((k2_filter_2(X0, X1) = k18_filter_2(k1_lattice2(X0), k5_filter_2(X0, X1))) & (k18_filter_2(X0, X1) = k2_filter_2(k1_lattice2(X0), k5_filter_2(X0, X1)))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f254])).
fof(f254, plain, ! [X0] : (! [X1] : (((k2_filter_2(X0, X1) = k18_filter_2(k1_lattice2(X0), k5_filter_2(X0, X1))) & (k18_filter_2(X0, X1) = k2_filter_2(k1_lattice2(X0), k5_filter_2(X0, X1)))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f111])).
fof(f111, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ((k2_filter_2(X0, X1) = k18_filter_2(k1_lattice2(X0), k5_filter_2(X0, X1))) & (k18_filter_2(X0, X1) = k2_filter_2(k1_lattice2(X0), k5_filter_2(X0, X1)))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', t30_filter_2)).
fof(f980, plain, ((k2_filter_0(k1_lattice2(sK1), sK2) = k2_filter_2(k1_lattice2(sK1), sK2)) | ~ spl33_22), inference(avatar_component_clause, [], [f978])).
fof(f508, plain, ! [X2, X0, X1] : (~ (u1_struct_0(X0) = k2_filter_0(X0, X1)) | (k4_lattices(X0, X1, X2) = X1) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f250])).
fof(f250, plain, ! [X0] : (! [X1] : (~ (u1_struct_0(X0) = k2_filter_0(X0, X1)) | ! [X2] : ((k4_lattices(X0, X1, X2) = X1) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f249])).
fof(f249, plain, ! [X0] : (! [X1] : ((~ (u1_struct_0(X0) = k2_filter_0(X0, X1)) | ! [X2] : ((k4_lattices(X0, X1, X2) = X1) | ~ m1_subset_1(X2, u1_struct_0(X0)))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f108])).
fof(f108, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ~ ((u1_struct_0(X0) = k2_filter_0(X0, X1)) & ~ ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (k4_lattices(X0, X1, X2) = X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', t23_filter_0)).
fof(f20611, plain, (spl33_92 | ~ spl33_80), inference(avatar_split_clause, [], [f20514, f2083, f2781])).
fof(f2083, plain, (spl33_80 <=> ! [X16, X15, X17] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15))), introduced(avatar_definition, [new_symbols(naming, [spl33_80])])).
fof(f20514, plain, ((u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | ~ spl33_80), inference(trivial_inequality_removal, [], [f20501])).
fof(f20501, plain, (~ (k1_lattice2(sK1) = k1_lattice2(sK1)) | (u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | ~ spl33_80), inference(superposition, [], [f2084, f3002])).
fof(f2084, plain, (! [X17, X15, X16] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15)) | ~ spl33_80), inference(avatar_component_clause, [], [f2083])).
fof(f15501, plain, spl33_77, inference(avatar_contradiction_clause, [], [f15500])).
fof(f15500, plain, ($false | spl33_77), inference(subsumption_resolution, [], [f15499, f540])).
fof(f15499, plain, (~ l2_lattices(sK1) | spl33_77), inference(resolution, [], [f2073, f788])).
fof(f788, plain, ! [X1] : (m1_relset_1(u2_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l2_lattices(X1)), inference(resolution, [], [f400, f501])).
fof(f501, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f331])).
fof(f331, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f103])).
fof(f103, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', redefinition_m2_relset_1)).
fof(f400, plain, ! [X0] : (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f196])).
fof(f196, plain, ! [X0] : ((m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0))) | ~ l2_lattices(X0)), inference(ennf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : (l2_lattices(X0) => (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', dt_u2_lattices)).
fof(f2073, plain, (~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl33_77), inference(avatar_component_clause, [], [f2071])).
fof(f2071, plain, (spl33_77 <=> m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl33_77])])).
fof(f15472, plain, spl33_76, inference(avatar_contradiction_clause, [], [f15471])).
fof(f15471, plain, ($false | spl33_76), inference(subsumption_resolution, [], [f15470, f540])).
fof(f15470, plain, (~ l2_lattices(sK1) | spl33_76), inference(resolution, [], [f2069, f399])).
fof(f399, plain, ! [X0] : (v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f196])).
fof(f2069, plain, (~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl33_76), inference(avatar_component_clause, [], [f2067])).
fof(f2067, plain, (spl33_76 <=> v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl33_76])])).
fof(f15395, plain, spl33_75, inference(avatar_contradiction_clause, [], [f15394])).
fof(f15394, plain, ($false | spl33_75), inference(subsumption_resolution, [], [f15393, f540])).
fof(f15393, plain, (~ l2_lattices(sK1) | spl33_75), inference(resolution, [], [f2065, f398])).
fof(f398, plain, ! [X0] : (v1_funct_1(u2_lattices(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f196])).
fof(f2065, plain, (~ v1_funct_1(u2_lattices(sK1)) | spl33_75), inference(avatar_component_clause, [], [f2063])).
fof(f2063, plain, (spl33_75 <=> v1_funct_1(u2_lattices(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl33_75])])).
fof(f15392, plain, spl33_74, inference(avatar_contradiction_clause, [], [f15391])).
fof(f15391, plain, ($false | spl33_74), inference(subsumption_resolution, [], [f15390, f534])).
fof(f534, plain, l1_lattices(sK1), inference(resolution, [], [f385, f338])).
fof(f15390, plain, (~ l1_lattices(sK1) | spl33_74), inference(resolution, [], [f2061, f786])).
fof(f786, plain, ! [X1] : (m1_relset_1(u1_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l1_lattices(X1)), inference(resolution, [], [f397, f501])).
fof(f397, plain, ! [X0] : (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f195])).
fof(f195, plain, ! [X0] : ((m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0))) | ~ l1_lattices(X0)), inference(ennf_transformation, [], [f52])).
fof(f52, plain, ! [X0] : (l1_lattices(X0) => (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', dt_u1_lattices)).
fof(f2061, plain, (~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl33_74), inference(avatar_component_clause, [], [f2059])).
fof(f2059, plain, (spl33_74 <=> m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl33_74])])).
fof(f15389, plain, (~ spl33_73 | ~ spl33_74 | ~ spl33_75 | ~ spl33_76 | ~ spl33_77 | spl33_78 | ~ spl33_72), inference(avatar_split_clause, [], [f15388, f2051, f2075, f2071, f2067, f2063, f2059, f2055])).
fof(f2055, plain, (spl33_73 <=> v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl33_73])])).
fof(f2051, plain, (spl33_72 <=> v1_funct_1(u1_lattices(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl33_72])])).
fof(f15388, plain, (! [X4, X5, X3] : (~ (g3_lattices(X3, X4, X5) = k1_lattice2(sK1)) | (u2_lattices(sK1) = X5) | ~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u2_lattices(sK1)) | ~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))) | ~ spl33_72), inference(subsumption_resolution, [], [f6185, f2052])).
fof(f2052, plain, (v1_funct_1(u1_lattices(sK1)) | ~ spl33_72), inference(avatar_component_clause, [], [f2051])).
fof(f6185, plain, ! [X4, X5, X3] : (~ (g3_lattices(X3, X4, X5) = k1_lattice2(sK1)) | (u2_lattices(sK1) = X5) | ~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u2_lattices(sK1)) | ~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u1_lattices(sK1))), inference(superposition, [], [f456, f768])).
fof(f768, plain, (k1_lattice2(sK1) = g3_lattices(u1_struct_0(sK1), u1_lattices(sK1), u2_lattices(sK1))), inference(resolution, [], [f364, f338])).
fof(f364, plain, ! [X0] : (~ l3_lattices(X0) | (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ! [X0] : ((k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : (l3_lattices(X0) => (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', d2_lattice2)).
fof(f456, plain, ! [X4, X2, X0, X5, X3, X1] : (~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) | (X2 = X5) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f228])).
fof(f228, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(flattening, [], [f227])).
fof(f227, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | (~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1))), inference(ennf_transformation, [], [f82])).
fof(f82, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X2) & m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X1)) => ! [X3, X4, X5] : ((g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) => ((X2 = X5) & (X1 = X4) & (X0 = X3)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', free_g3_lattices)).
fof(f15385, plain, (~ spl33_73 | ~ spl33_74 | ~ spl33_75 | ~ spl33_76 | ~ spl33_77 | spl33_80 | ~ spl33_72), inference(avatar_split_clause, [], [f15384, f2051, f2083, f2071, f2067, f2063, f2059, f2055])).
fof(f15384, plain, (! [X17, X15, X16] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15) | ~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u2_lattices(sK1)) | ~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))) | ~ spl33_72), inference(subsumption_resolution, [], [f6189, f2052])).
fof(f6189, plain, ! [X17, X15, X16] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15) | ~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u2_lattices(sK1)) | ~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u1_lattices(sK1))), inference(superposition, [], [f454, f768])).
fof(f454, plain, ! [X4, X2, X0, X5, X3, X1] : (~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) | (X0 = X3) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f228])).
fof(f15383, plain, spl33_73, inference(avatar_contradiction_clause, [], [f15382])).
fof(f15382, plain, ($false | spl33_73), inference(subsumption_resolution, [], [f15381, f534])).
fof(f15381, plain, (~ l1_lattices(sK1) | spl33_73), inference(resolution, [], [f2057, f396])).
fof(f396, plain, ! [X0] : (v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f195])).
fof(f2057, plain, (~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl33_73), inference(avatar_component_clause, [], [f2055])).
fof(f2088, plain, spl33_72, inference(avatar_contradiction_clause, [], [f2087])).
fof(f2087, plain, ($false | spl33_72), inference(subsumption_resolution, [], [f2086, f534])).
fof(f2086, plain, (~ l1_lattices(sK1) | spl33_72), inference(resolution, [], [f2053, f395])).
fof(f395, plain, ! [X0] : (v1_funct_1(u1_lattices(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f195])).
fof(f2053, plain, (~ v1_funct_1(u1_lattices(sK1)) | spl33_72), inference(avatar_component_clause, [], [f2051])).
fof(f1025, plain, ~ spl33_14, inference(avatar_contradiction_clause, [], [f1024])).
fof(f1024, plain, ($false | ~ spl33_14), inference(subsumption_resolution, [], [f1020, f671])).
fof(f671, plain, sP0(sK1), inference(subsumption_resolution, [], [f670, f336])).
fof(f670, plain, (sP0(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f665, f337])).
fof(f665, plain, (sP0(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f445, f338])).
fof(f445, plain, ! [X0] : (~ l3_lattices(X0) | sP0(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f264])).
fof(f264, plain, ! [X0] : (sP0(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(definition_folding, [], [f222, e263])).
fof(f263, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(usedef, [], [e263])).
fof(e263, plain, ! [X0] : (sP0(X0) <=> (v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f222, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f221])).
fof(f221, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f79])).
fof(f79, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', fc6_lattice2)).
fof(f1020, plain, (~ sP0(sK1) | ~ spl33_14), inference(resolution, [], [f928, f436])).
fof(f436, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f291])).
fof(f291, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(nnf_transformation, [], [f263])).
fof(f928, plain, (v3_struct_0(k1_lattice2(sK1)) | ~ spl33_14), inference(avatar_component_clause, [], [f926])).
fof(f1005, plain, spl33_16, inference(avatar_contradiction_clause, [], [f1004])).
fof(f1004, plain, ($false | spl33_16), inference(subsumption_resolution, [], [f1003, f338])).
fof(f1003, plain, (~ l3_lattices(sK1) | spl33_16), inference(resolution, [], [f936, f374])).
fof(f936, plain, (~ l3_lattices(k1_lattice2(sK1)) | spl33_16), inference(avatar_component_clause, [], [f934])).
fof(f981, plain, (spl33_14 | ~ spl33_15 | ~ spl33_16 | spl33_22), inference(avatar_split_clause, [], [f964, f978, f934, f930, f926])).
fof(f964, plain, ((k2_filter_0(k1_lattice2(sK1), sK2) = k2_filter_2(k1_lattice2(sK1), sK2)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(resolution, [], [f496, f916])).
fof(f496, plain, ! [X0, X1] : (~ m1_subset_1(X1, u1_struct_0(X0)) | (k2_filter_0(X0, X1) = k2_filter_2(X0, X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f237])).
fof(f237, plain, ! [X0, X1] : ((k2_filter_0(X0, X1) = k2_filter_2(X0, X1)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f236])).
fof(f236, plain, ! [X0, X1] : ((k2_filter_0(X0, X1) = k2_filter_2(X0, X1)) | (~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f99])).
fof(f99, plain, ! [X0, X1] : ((m1_subset_1(X1, u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k2_filter_0(X0, X1) = k2_filter_2(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT308+1.p', redefinition_k2_filter_2)).
fof(f976, plain, spl33_15, inference(avatar_contradiction_clause, [], [f975])).
fof(f975, plain, ($false | spl33_15), inference(subsumption_resolution, [], [f974, f671])).
fof(f974, plain, (~ sP0(sK1) | spl33_15), inference(resolution, [], [f932, f444])).
fof(f444, plain, ! [X0] : (v10_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f291])).
fof(f932, plain, (~ v10_lattices(k1_lattice2(sK1)) | spl33_15), inference(avatar_component_clause, [], [f930])).