fof(f14837, plain, $false, inference(avatar_sat_refutation, [], [f800, f1112, f1131, f14232, f14350, f14639, f14648, f14651, f14817, f14836])).
fof(f14836, plain, (spl30_2 | ~ spl30_415), inference(avatar_contradiction_clause, [], [f14835])).
fof(f14835, plain, ($false | (spl30_2 | ~ spl30_415)), inference(subsumption_resolution, [], [f14834, f415])).
fof(f415, plain, v2_orders_2(sK4), inference(cnf_transformation, [], [f357])).
fof(f357, plain, ((~ (k7_grcat_1(sK4) = k2_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4)))) & l1_orders_2(sK4) & v3_lattice3(sK4) & v2_lattice3(sK4) & v1_lattice3(sK4) & v4_orders_2(sK4) & v3_orders_2(sK4) & v2_orders_2(sK4)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4])], [f194, f356])).
fof(f356, plain, (? [X0] : ((~ (k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) | ~ (k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0)))) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ((~ (k7_grcat_1(sK4) = k2_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4)))) & l1_orders_2(sK4) & v3_lattice3(sK4) & v2_lattice3(sK4) & v1_lattice3(sK4) & v4_orders_2(sK4) & v3_orders_2(sK4) & v2_orders_2(sK4))), introduced(choice_axiom, [])).
fof(f194, plain, ? [X0] : ((~ (k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) | ~ (k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0)))) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(flattening, [], [f193])).
fof(f193, plain, ? [X0] : ((~ (k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) | ~ (k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0)))) & (l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ((k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) & (k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ((k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) & (k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', t7_waybel34)).
fof(f14834, plain, (~ v2_orders_2(sK4) | (spl30_2 | ~ spl30_415)), inference(subsumption_resolution, [], [f14833, f416])).
fof(f416, plain, v3_orders_2(sK4), inference(cnf_transformation, [], [f357])).
fof(f14833, plain, (~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (spl30_2 | ~ spl30_415)), inference(subsumption_resolution, [], [f14832, f417])).
fof(f417, plain, v4_orders_2(sK4), inference(cnf_transformation, [], [f357])).
fof(f14832, plain, (~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (spl30_2 | ~ spl30_415)), inference(subsumption_resolution, [], [f14831, f418])).
fof(f418, plain, v1_lattice3(sK4), inference(cnf_transformation, [], [f357])).
fof(f14831, plain, (~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (spl30_2 | ~ spl30_415)), inference(subsumption_resolution, [], [f14830, f419])).
fof(f419, plain, v2_lattice3(sK4), inference(cnf_transformation, [], [f357])).
fof(f14830, plain, (~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (spl30_2 | ~ spl30_415)), inference(subsumption_resolution, [], [f14829, f421])).
fof(f421, plain, l1_orders_2(sK4), inference(cnf_transformation, [], [f357])).
fof(f14829, plain, (~ l1_orders_2(sK4) | ~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (spl30_2 | ~ spl30_415)), inference(subsumption_resolution, [], [f14828, f420])).
fof(f420, plain, v3_lattice3(sK4), inference(cnf_transformation, [], [f357])).
fof(f14828, plain, (~ v3_lattice3(sK4) | ~ l1_orders_2(sK4) | ~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (spl30_2 | ~ spl30_415)), inference(subsumption_resolution, [], [f14826, f799])).
fof(f799, plain, (~ (k7_grcat_1(sK4) = k2_waybel34(sK4, sK4, k7_grcat_1(sK4))) | spl30_2), inference(avatar_component_clause, [], [f797])).
fof(f797, plain, (spl30_2 <=> (k7_grcat_1(sK4) = k2_waybel34(sK4, sK4, k7_grcat_1(sK4)))), introduced(avatar_definition, [new_symbols(naming, [spl30_2])])).
fof(f14826, plain, ((k7_grcat_1(sK4) = k2_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ v3_lattice3(sK4) | ~ l1_orders_2(sK4) | ~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | ~ spl30_415), inference(resolution, [], [f14349, f4130])).
fof(f4130, plain, ! [X0] : (~ v18_waybel_0(k7_grcat_1(X0), X0, X0) | (k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(subsumption_resolution, [], [f4129, f573])).
fof(f573, plain, ! [X0] : (~ l1_orders_2(X0) | l1_struct_0(X0)), inference(cnf_transformation, [], [f274])).
fof(f274, plain, ! [X0] : (l1_struct_0(X0) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f62])).
fof(f62, plain, ! [X0] : (l1_orders_2(X0) => l1_struct_0(X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', dt_l1_orders_2)).
fof(f4129, plain, ! [X0] : ((k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) | ~ v18_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | ~ l1_struct_0(X0)), inference(resolution, [], [f2647, f572])).
fof(f572, plain, ! [X0] : (m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ l1_struct_0(X0)), inference(cnf_transformation, [], [f273])).
fof(f273, plain, ! [X0] : ((m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(k7_grcat_1(X0))) | ~ l1_struct_0(X0)), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ! [X0] : (l1_struct_0(X0) => (m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(k7_grcat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', dt_k7_grcat_1)).
fof(f2647, plain, ! [X0] : (~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | (k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) | ~ v18_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(subsumption_resolution, [], [f2646, f627])).
fof(f627, plain, ! [X0] : (v1_funct_1(k7_grcat_1(X0)) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f297])).
fof(f297, plain, ! [X0] : ((v23_waybel_0(k7_grcat_1(X0), X0, X0) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0))) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f89])).
fof(f89, plain, ! [X0] : (l1_orders_2(X0) => (v23_waybel_0(k7_grcat_1(X0), X0, X0) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', fc2_waybel_1)).
fof(f2646, plain, ! [X0] : ((k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) | ~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(k7_grcat_1(X0)) | ~ v18_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(subsumption_resolution, [], [f2645, f628])).
fof(f628, plain, ! [X0] : (v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f297])).
fof(f2645, plain, ! [X0] : ((k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) | ~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(k7_grcat_1(X0)) | ~ v18_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(subsumption_resolution, [], [f2644, f483])).
fof(f483, plain, ! [X0] : (~ v3_struct_0(X0) | ~ v1_lattice3(X0) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f219])).
fof(f219, plain, ! [X0] : (~ v3_struct_0(X0) | ~ v1_lattice3(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f218])).
fof(f218, plain, ! [X0] : ((~ v3_struct_0(X0) | ~ v1_lattice3(X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ! [X0] : (l1_orders_2(X0) => (v1_lattice3(X0) => ~ v3_struct_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', cc1_lattice3)).
fof(f2644, plain, ! [X0] : ((k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) | ~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(k7_grcat_1(X0)) | ~ v18_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(duplicate_literal_removal, [], [f2643])).
fof(f2643, plain, ! [X0] : ((k7_grcat_1(X0) = k2_waybel34(X0, X0, k7_grcat_1(X0))) | ~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(k7_grcat_1(X0)) | ~ v18_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ v3_lattice3(X0) | ~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(k7_grcat_1(X0)) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(resolution, [], [f562, f786])).
fof(f786, plain, ! [X0] : (v3_waybel_1(k1_waybel_1(X0, X0, k7_grcat_1(X0), k7_grcat_1(X0)), X0, X0) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f345])).
fof(f345, plain, ! [X0] : (v3_waybel_1(k1_waybel_1(X0, X0, k7_grcat_1(X0), k7_grcat_1(X0)), X0, X0) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f344])).
fof(f344, plain, ! [X0] : (v3_waybel_1(k1_waybel_1(X0, X0, k7_grcat_1(X0), k7_grcat_1(X0)), X0, X0) | (~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f129])).
fof(f129, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => v3_waybel_1(k1_waybel_1(X0, X0, k7_grcat_1(X0), k7_grcat_1(X0)), X0, X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', t6_waybel34)).
fof(f562, plain, ! [X2, X0, X3, X1] : (~ v3_waybel_1(k1_waybel_1(X0, X1, X3, X2), X0, X1) | (k2_waybel34(X0, X1, X2) = X3) | ~ m2_relset_1(X3, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X3, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X3) | ~ v18_waybel_0(X2, X1, X0) | ~ v3_lattice3(X1) | ~ v3_lattice3(X0) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f361])).
fof(f361, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((((k2_waybel34(X0, X1, X2) = X3) | ~ v3_waybel_1(k1_waybel_1(X0, X1, X3, X2), X0, X1)) & (v3_waybel_1(k1_waybel_1(X0, X1, X3, X2), X0, X1) | ~ (k2_waybel34(X0, X1, X2) = X3))) | ~ m2_relset_1(X3, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X3, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X3)) | ~ v18_waybel_0(X2, X1, X0) | ~ v3_lattice3(X1) | ~ v3_lattice3(X0) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(nnf_transformation, [], [f268])).
fof(f268, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((k2_waybel34(X0, X1, X2) = X3) <=> v3_waybel_1(k1_waybel_1(X0, X1, X3, X2), X0, X1)) | ~ m2_relset_1(X3, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X3, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X3)) | ~ v18_waybel_0(X2, X1, X0) | ~ v3_lattice3(X1) | ~ v3_lattice3(X0) | ~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f267])).
fof(f267, plain, ! [X0] : (! [X1] : (! [X2] : ((! [X3] : (((k2_waybel34(X0, X1, X2) = X3) <=> v3_waybel_1(k1_waybel_1(X0, X1, X3, X2), X0, X1)) | (~ m2_relset_1(X3, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X3, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X3))) | (~ v18_waybel_0(X2, X1, X0) | ~ v3_lattice3(X1) | ~ v3_lattice3(X0))) | (~ m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f50])).
fof(f50, plain, ! [X0] : ((l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2)) => ((v18_waybel_0(X2, X1, X0) & v3_lattice3(X1) & v3_lattice3(X0)) => ! [X3] : ((m2_relset_1(X3, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X3, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X3)) => ((k2_waybel34(X0, X1, X2) = X3) <=> v3_waybel_1(k1_waybel_1(X0, X1, X3, X2), X0, X1))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', d2_waybel34)).
fof(f14349, plain, (v18_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ spl30_415), inference(avatar_component_clause, [], [f14347])).
fof(f14347, plain, (spl30_415 <=> v18_waybel_0(k7_grcat_1(sK4), sK4, sK4)), introduced(avatar_definition, [new_symbols(naming, [spl30_415])])).
fof(f14817, plain, spl30_414, inference(avatar_contradiction_clause, [], [f14816])).
fof(f14816, plain, ($false | spl30_414), inference(subsumption_resolution, [], [f14815, f834])).
fof(f834, plain, l1_struct_0(sK4), inference(resolution, [], [f573, f421])).
fof(f14815, plain, (~ l1_struct_0(sK4) | spl30_414), inference(resolution, [], [f14345, f1150])).
fof(f1150, plain, ! [X0] : (m1_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ l1_struct_0(X0)), inference(resolution, [], [f572, f777])).
fof(f777, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f414])).
fof(f414, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f121])).
fof(f121, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', redefinition_m2_relset_1)).
fof(f14345, plain, (~ m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | spl30_414), inference(avatar_component_clause, [], [f14343])).
fof(f14343, plain, (spl30_414 <=> m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4))), introduced(avatar_definition, [new_symbols(naming, [spl30_414])])).
fof(f14651, plain, (~ spl30_8 | spl30_413), inference(avatar_contradiction_clause, [], [f14650])).
fof(f14650, plain, ($false | (~ spl30_8 | spl30_413)), inference(subsumption_resolution, [], [f14649, f1111])).
fof(f1111, plain, (sP2(sK4) | ~ spl30_8), inference(avatar_component_clause, [], [f1109])).
fof(f1109, plain, (spl30_8 <=> sP2(sK4)), introduced(avatar_definition, [new_symbols(naming, [spl30_8])])).
fof(f14649, plain, (~ sP2(sK4) | spl30_413), inference(resolution, [], [f14341, f645])).
fof(f645, plain, ! [X0] : (v17_waybel_0(k7_grcat_1(X0), X0, X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f372])).
fof(f372, plain, ! [X0] : ((v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0))) | ~ sP2(X0)), inference(nnf_transformation, [], [f352])).
fof(f352, plain, ! [X0] : ((v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0))) | ~ sP2(X0)), inference(usedef, [], [e352])).
fof(e352, plain, ! [X0] : (sP2(X0) <=> (v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f14341, plain, (~ v17_waybel_0(k7_grcat_1(sK4), sK4, sK4) | spl30_413), inference(avatar_component_clause, [], [f14339])).
fof(f14339, plain, (spl30_413 <=> v17_waybel_0(k7_grcat_1(sK4), sK4, sK4)), introduced(avatar_definition, [new_symbols(naming, [spl30_413])])).
fof(f14648, plain, spl30_411, inference(avatar_contradiction_clause, [], [f14647])).
fof(f14647, plain, ($false | spl30_411), inference(subsumption_resolution, [], [f14642, f834])).
fof(f14642, plain, (~ l1_struct_0(sK4) | spl30_411), inference(resolution, [], [f14333, f570])).
fof(f570, plain, ! [X0] : (v1_funct_1(k7_grcat_1(X0)) | ~ l1_struct_0(X0)), inference(cnf_transformation, [], [f273])).
fof(f14333, plain, (~ v1_funct_1(k7_grcat_1(sK4)) | spl30_411), inference(avatar_component_clause, [], [f14331])).
fof(f14331, plain, (spl30_411 <=> v1_funct_1(k7_grcat_1(sK4))), introduced(avatar_definition, [new_symbols(naming, [spl30_411])])).
fof(f14639, plain, spl30_412, inference(avatar_contradiction_clause, [], [f14638])).
fof(f14638, plain, ($false | spl30_412), inference(subsumption_resolution, [], [f14633, f834])).
fof(f14633, plain, (~ l1_struct_0(sK4) | spl30_412), inference(resolution, [], [f14337, f571])).
fof(f571, plain, ! [X0] : (v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ l1_struct_0(X0)), inference(cnf_transformation, [], [f273])).
fof(f14337, plain, (~ v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | spl30_412), inference(avatar_component_clause, [], [f14335])).
fof(f14335, plain, (spl30_412 <=> v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4))), introduced(avatar_definition, [new_symbols(naming, [spl30_412])])).
fof(f14350, plain, (~ spl30_411 | ~ spl30_412 | ~ spl30_413 | ~ spl30_414 | spl30_415 | ~ spl30_1), inference(avatar_split_clause, [], [f14329, f793, f14347, f14343, f14339, f14335, f14331])).
fof(f793, plain, (spl30_1 <=> (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4)))), introduced(avatar_definition, [new_symbols(naming, [spl30_1])])).
fof(f14329, plain, (v18_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v17_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v1_funct_1(k7_grcat_1(sK4)) | ~ spl30_1), inference(subsumption_resolution, [], [f14328, f415])).
fof(f14328, plain, (v18_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v17_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v1_funct_1(k7_grcat_1(sK4)) | ~ v2_orders_2(sK4) | ~ spl30_1), inference(subsumption_resolution, [], [f14327, f416])).
fof(f14327, plain, (v18_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v17_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v1_funct_1(k7_grcat_1(sK4)) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | ~ spl30_1), inference(subsumption_resolution, [], [f14326, f417])).
fof(f14326, plain, (v18_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v17_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v1_funct_1(k7_grcat_1(sK4)) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | ~ spl30_1), inference(subsumption_resolution, [], [f14325, f418])).
fof(f14325, plain, (v18_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v17_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v1_funct_1(k7_grcat_1(sK4)) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | ~ spl30_1), inference(subsumption_resolution, [], [f14324, f419])).
fof(f14324, plain, (v18_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v17_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v1_funct_1(k7_grcat_1(sK4)) | ~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | ~ spl30_1), inference(subsumption_resolution, [], [f14323, f420])).
fof(f14323, plain, (v18_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v17_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v1_funct_1(k7_grcat_1(sK4)) | ~ v3_lattice3(sK4) | ~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | ~ spl30_1), inference(subsumption_resolution, [], [f14320, f421])).
fof(f14320, plain, (v18_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v17_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v1_funct_1(k7_grcat_1(sK4)) | ~ l1_orders_2(sK4) | ~ v3_lattice3(sK4) | ~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | ~ spl30_1), inference(duplicate_literal_removal, [], [f14235])).
fof(f14235, plain, (v18_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ m1_relset_1(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v17_waybel_0(k7_grcat_1(sK4), sK4, sK4) | ~ v1_funct_2(k7_grcat_1(sK4), u1_struct_0(sK4), u1_struct_0(sK4)) | ~ v1_funct_1(k7_grcat_1(sK4)) | ~ l1_orders_2(sK4) | ~ v3_lattice3(sK4) | ~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | ~ l1_orders_2(sK4) | ~ v3_lattice3(sK4) | ~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | ~ spl30_1), inference(superposition, [], [f612, f794])).
fof(f794, plain, ((k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ spl30_1), inference(avatar_component_clause, [], [f793])).
fof(f612, plain, ! [X2, X0, X1] : (v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f293])).
fof(f293, plain, ! [X0, X1, X2] : ((v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f292])).
fof(f292, plain, ! [X0, X1, X2] : ((v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f187])).
fof(f187, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f164])).
fof(f164, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_partfun1(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f140])).
fof(f140, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k1_waybel34(X0, X1, X2), X1, X0) & v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_partfun1(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f85])).
fof(f85, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k1_waybel34(X0, X1, X2), X1, X0) & v5_waybel_1(k1_waybel34(X0, X1, X2), X0, X1) & v22_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v20_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v18_waybel_0(k1_waybel34(X0, X1, X2), X1, X0) & v1_partfun1(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(k1_waybel34(X0, X1, X2), u1_struct_0(X1), u1_struct_0(X0)) & ~ v1_xboole_0(k1_waybel34(X0, X1, X2)) & v1_funct_1(k1_waybel34(X0, X1, X2)) & v1_relat_1(k1_waybel34(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', fc1_waybel34)).
fof(f14232, plain, (spl30_1 | ~ spl30_8), inference(avatar_split_clause, [], [f14222, f1109, f793])).
fof(f14222, plain, ((k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ spl30_8), inference(subsumption_resolution, [], [f14221, f415])).
fof(f14221, plain, (~ v2_orders_2(sK4) | (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ spl30_8), inference(subsumption_resolution, [], [f14220, f416])).
fof(f14220, plain, (~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ spl30_8), inference(subsumption_resolution, [], [f14219, f417])).
fof(f14219, plain, (~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ spl30_8), inference(subsumption_resolution, [], [f14218, f418])).
fof(f14218, plain, (~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ spl30_8), inference(subsumption_resolution, [], [f14217, f419])).
fof(f14217, plain, (~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ spl30_8), inference(subsumption_resolution, [], [f14216, f421])).
fof(f14216, plain, (~ l1_orders_2(sK4) | ~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ spl30_8), inference(subsumption_resolution, [], [f14214, f420])).
fof(f14214, plain, (~ v3_lattice3(sK4) | ~ l1_orders_2(sK4) | ~ v2_lattice3(sK4) | ~ v1_lattice3(sK4) | ~ v4_orders_2(sK4) | ~ v3_orders_2(sK4) | ~ v2_orders_2(sK4) | (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ spl30_8), inference(resolution, [], [f6371, f1111])).
fof(f6371, plain, ! [X0] : (~ sP2(X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | (k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0)))), inference(resolution, [], [f4128, f645])).
fof(f4128, plain, ! [X0] : (~ v17_waybel_0(k7_grcat_1(X0), X0, X0) | (k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0))) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(subsumption_resolution, [], [f4127, f573])).
fof(f4127, plain, ! [X0] : ((k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0))) | ~ v17_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | ~ l1_struct_0(X0)), inference(resolution, [], [f2641, f572])).
fof(f2641, plain, ! [X0] : (~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | (k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0))) | ~ v17_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(subsumption_resolution, [], [f2640, f627])).
fof(f2640, plain, ! [X0] : ((k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0))) | ~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(k7_grcat_1(X0)) | ~ v17_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(subsumption_resolution, [], [f2639, f628])).
fof(f2639, plain, ! [X0] : ((k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0))) | ~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(k7_grcat_1(X0)) | ~ v17_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(subsumption_resolution, [], [f2638, f483])).
fof(f2638, plain, ! [X0] : ((k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0))) | ~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(k7_grcat_1(X0)) | ~ v17_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(duplicate_literal_removal, [], [f2637])).
fof(f2637, plain, ! [X0] : ((k7_grcat_1(X0) = k1_waybel34(X0, X0, k7_grcat_1(X0))) | ~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(k7_grcat_1(X0)) | ~ v17_waybel_0(k7_grcat_1(X0), X0, X0) | ~ v3_lattice3(X0) | ~ v3_lattice3(X0) | ~ m2_relset_1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(k7_grcat_1(X0)) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(resolution, [], [f560, f786])).
fof(f560, plain, ! [X2, X0, X3, X1] : (~ v3_waybel_1(k1_waybel_1(X0, X1, X2, X3), X0, X1) | (k1_waybel34(X0, X1, X2) = X3) | ~ m2_relset_1(X3, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X3, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X3) | ~ v17_waybel_0(X2, X0, X1) | ~ v3_lattice3(X1) | ~ v3_lattice3(X0) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f360])).
fof(f360, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((((k1_waybel34(X0, X1, X2) = X3) | ~ v3_waybel_1(k1_waybel_1(X0, X1, X2, X3), X0, X1)) & (v3_waybel_1(k1_waybel_1(X0, X1, X2, X3), X0, X1) | ~ (k1_waybel34(X0, X1, X2) = X3))) | ~ m2_relset_1(X3, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X3, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X3)) | ~ v17_waybel_0(X2, X0, X1) | ~ v3_lattice3(X1) | ~ v3_lattice3(X0) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(nnf_transformation, [], [f266])).
fof(f266, plain, ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((k1_waybel34(X0, X1, X2) = X3) <=> v3_waybel_1(k1_waybel_1(X0, X1, X2, X3), X0, X1)) | ~ m2_relset_1(X3, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X3, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X3)) | ~ v17_waybel_0(X2, X0, X1) | ~ v3_lattice3(X1) | ~ v3_lattice3(X0) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f265])).
fof(f265, plain, ! [X0] : (! [X1] : (! [X2] : ((! [X3] : (((k1_waybel34(X0, X1, X2) = X3) <=> v3_waybel_1(k1_waybel_1(X0, X1, X2, X3), X0, X1)) | (~ m2_relset_1(X3, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X3, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X3))) | (~ v17_waybel_0(X2, X0, X1) | ~ v3_lattice3(X1) | ~ v3_lattice3(X0))) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f49])).
fof(f49, plain, ! [X0] : ((l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => ((v17_waybel_0(X2, X0, X1) & v3_lattice3(X1) & v3_lattice3(X0)) => ! [X3] : ((m2_relset_1(X3, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(X3, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X3)) => ((k1_waybel34(X0, X1, X2) = X3) <=> v3_waybel_1(k1_waybel_1(X0, X1, X2, X3), X0, X1))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', d1_waybel34)).
fof(f1131, plain, ~ spl30_7, inference(avatar_contradiction_clause, [], [f1130])).
fof(f1130, plain, ($false | ~ spl30_7), inference(subsumption_resolution, [], [f1129, f421])).
fof(f1129, plain, (~ l1_orders_2(sK4) | ~ spl30_7), inference(subsumption_resolution, [], [f1125, f418])).
fof(f1125, plain, (~ v1_lattice3(sK4) | ~ l1_orders_2(sK4) | ~ spl30_7), inference(resolution, [], [f1033, f483])).
fof(f1033, plain, (v3_struct_0(sK4) | ~ spl30_7), inference(avatar_component_clause, [], [f1031])).
fof(f1031, plain, (spl30_7 <=> v3_struct_0(sK4)), introduced(avatar_definition, [new_symbols(naming, [spl30_7])])).
fof(f1112, plain, (spl30_7 | spl30_8), inference(avatar_split_clause, [], [f1107, f1109, f1031])).
fof(f1107, plain, (sP2(sK4) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f1104, f415])).
fof(f1104, plain, (sP2(sK4) | ~ v2_orders_2(sK4) | v3_struct_0(sK4)), inference(resolution, [], [f653, f421])).
fof(f653, plain, ! [X0] : (~ l1_orders_2(X0) | sP2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f353])).
fof(f353, plain, ! [X0] : (sP2(X0) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(definition_folding, [], [f304, e352])).
fof(f304, plain, ! [X0] : ((v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0))) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f303])).
fof(f303, plain, ! [X0] : ((v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0))) | (~ l1_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f190])).
fof(f190, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0)))), inference(pure_predicate_removal, [], [f172])).
fof(f172, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v1_partfun1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0)))), inference(pure_predicate_removal, [], [f162])).
fof(f162, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v11_quantal1(k7_grcat_1(X0)) & v1_partfun1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0)))), inference(pure_predicate_removal, [], [f156])).
fof(f156, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v5_orders_3(k7_grcat_1(X0), X0, X0) & v11_quantal1(k7_grcat_1(X0)) & v1_partfun1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0)))), inference(pure_predicate_removal, [], [f151])).
fof(f151, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v21_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v5_orders_3(k7_grcat_1(X0), X0, X0) & v11_quantal1(k7_grcat_1(X0)) & v1_partfun1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0)))), inference(pure_predicate_removal, [], [f141])).
fof(f141, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v21_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v5_orders_3(k7_grcat_1(X0), X0, X0) & v11_quantal1(k7_grcat_1(X0)) & v1_partfun1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v2_funct_1(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0)))), inference(pure_predicate_removal, [], [f93])).
fof(f93, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v8_waybel_1(k7_grcat_1(X0), X0) & v7_waybel_1(k7_grcat_1(X0), X0) & v6_waybel_1(k7_grcat_1(X0), X0) & v23_waybel_0(k7_grcat_1(X0), X0, X0) & v22_waybel_0(k7_grcat_1(X0), X0, X0) & v21_waybel_0(k7_grcat_1(X0), X0, X0) & v20_waybel_0(k7_grcat_1(X0), X0, X0) & v19_waybel_0(k7_grcat_1(X0), X0, X0) & v18_waybel_0(k7_grcat_1(X0), X0, X0) & v17_waybel_0(k7_grcat_1(X0), X0, X0) & v5_orders_3(k7_grcat_1(X0), X0, X0) & v11_quantal1(k7_grcat_1(X0)) & v1_partfun1(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(k7_grcat_1(X0), u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(k7_grcat_1(X0)) & v2_funct_1(k7_grcat_1(X0)) & v1_funct_1(k7_grcat_1(X0)) & v1_relat_1(k7_grcat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT356+1.p', fc4_waybel10)).
fof(f800, plain, (~ spl30_1 | ~ spl30_2), inference(avatar_split_clause, [], [f422, f797, f793])).
fof(f422, plain, (~ (k7_grcat_1(sK4) = k2_waybel34(sK4, sK4, k7_grcat_1(sK4))) | ~ (k7_grcat_1(sK4) = k1_waybel34(sK4, sK4, k7_grcat_1(sK4)))), inference(cnf_transformation, [], [f357])).