fof(f74410, plain, $false, inference(avatar_sat_refutation, [], [f566, f567, f568, f1554, f1561, f1566, f1570, f1698, f2039, f2045, f2051, f2054, f2263, f2373, f2376, f2542, f4780, f4781, f4830, f4835, f68484, f68485, f68737, f69356, f74409])).
fof(f74409, plain, (spl32_1 | ~ spl32_31 | ~ spl32_33 | ~ spl32_45 | ~ spl32_46 | ~ spl32_109 | ~ spl32_2724 | ~ spl32_2725), inference(avatar_contradiction_clause, [], [f74408])).
fof(f74408, plain, ($false | (spl32_1 | ~ spl32_31 | ~ spl32_33 | ~ spl32_45 | ~ spl32_46 | ~ spl32_109 | ~ spl32_2724 | ~ spl32_2725)), inference(subsumption_resolution, [], [f74407, f556])).
fof(f556, plain, (~ (k17_filter_2(sK2) = sK3) | spl32_1), inference(avatar_component_clause, [], [f555])).
fof(f555, plain, (spl32_1 <=> (k17_filter_2(sK2) = sK3)), introduced(avatar_definition, [new_symbols(naming, [spl32_1])])).
fof(f74407, plain, ((k17_filter_2(sK2) = sK3) | (~ spl32_31 | ~ spl32_33 | ~ spl32_45 | ~ spl32_46 | ~ spl32_109 | ~ spl32_2724 | ~ spl32_2725)), inference(trivial_inequality_removal, [], [f74355])).
fof(f74355, plain, (~ (k1_lattice2(sK2) = k1_lattice2(sK2)) | (k17_filter_2(sK2) = sK3) | (~ spl32_31 | ~ spl32_33 | ~ spl32_45 | ~ spl32_46 | ~ spl32_109 | ~ spl32_2724 | ~ spl32_2725)), inference(superposition, [], [f2050, f69029])).
fof(f69029, plain, ((k1_lattice2(sK2) = g3_lattices(sK3, u1_lattices(sK2), u2_lattices(sK2))) | (~ spl32_31 | ~ spl32_33 | ~ spl32_45 | ~ spl32_46 | ~ spl32_2724 | ~ spl32_2725)), inference(backward_demodulation, [], [f68847, f68848])).
fof(f68848, plain, ((sK3 = k17_filter_2(k1_lattice2(sK2))) | (~ spl32_33 | ~ spl32_45 | ~ spl32_46)), inference(forward_demodulation, [], [f1700, f68819])).
fof(f68819, plain, ((sK3 = k1_filter_0(k1_lattice2(sK2))) | ~ spl32_33), inference(forward_demodulation, [], [f1553, f2660])).
fof(f2660, plain, (sK3 = k7_filter_2(sK2, sK3)), inference(subsumption_resolution, [], [f2659, f305])).
fof(f305, plain, ~ v3_struct_0(sK2), inference(cnf_transformation, [], [f240])).
fof(f240, plain, (((~ r2_filter_2(sK2, sK3) | ~ v1_filter_2(sK3, sK2) | (k17_filter_2(sK2) = sK3)) & (r2_filter_2(sK2, sK3) | (v1_filter_2(sK3, sK2) & ~ (k17_filter_2(sK2) = sK3))) & m2_filter_2(sK3, sK2)) & l3_lattices(sK2) & v17_lattices(sK2) & v10_lattices(sK2) & ~ v3_struct_0(sK2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2, sK3])], [f237, f239, f238])).
fof(f238, plain, (? [X0] : (? [X1] : ((~ r2_filter_2(X0, X1) | ~ v1_filter_2(X1, X0) | (k17_filter_2(X0) = X1)) & (r2_filter_2(X0, X1) | (v1_filter_2(X1, X0) & ~ (k17_filter_2(X0) = X1))) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (? [X1] : ((~ r2_filter_2(sK2, X1) | ~ v1_filter_2(X1, sK2) | (k17_filter_2(sK2) = X1)) & (r2_filter_2(sK2, X1) | (v1_filter_2(X1, sK2) & ~ (k17_filter_2(sK2) = X1))) & m2_filter_2(X1, sK2)) & l3_lattices(sK2) & v17_lattices(sK2) & v10_lattices(sK2) & ~ v3_struct_0(sK2))), introduced(choice_axiom, [])).
fof(f239, plain, (? [X1] : ((~ r2_filter_2(sK2, X1) | ~ v1_filter_2(X1, sK2) | (k17_filter_2(sK2) = X1)) & (r2_filter_2(sK2, X1) | (v1_filter_2(X1, sK2) & ~ (k17_filter_2(sK2) = X1))) & m2_filter_2(X1, sK2)) => ((~ r2_filter_2(sK2, sK3) | ~ v1_filter_2(sK3, sK2) | (k17_filter_2(sK2) = sK3)) & (r2_filter_2(sK2, sK3) | (v1_filter_2(sK3, sK2) & ~ (k17_filter_2(sK2) = sK3))) & m2_filter_2(sK3, sK2))), introduced(choice_axiom, [])).
fof(f237, plain, ? [X0] : (? [X1] : ((~ r2_filter_2(X0, X1) | ~ v1_filter_2(X1, X0) | (k17_filter_2(X0) = X1)) & (r2_filter_2(X0, X1) | (v1_filter_2(X1, X0) & ~ (k17_filter_2(X0) = X1))) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f236])).
fof(f236, plain, ? [X0] : (? [X1] : (((~ r2_filter_2(X0, X1) | (~ v1_filter_2(X1, X0) | (k17_filter_2(X0) = X1))) & (r2_filter_2(X0, X1) | (v1_filter_2(X1, X0) & ~ (k17_filter_2(X0) = X1)))) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(nnf_transformation, [], [f117])).
fof(f117, plain, ? [X0] : (? [X1] : (~ ((v1_filter_2(X1, X0) & ~ (k17_filter_2(X0) = X1)) <=> r2_filter_2(X0, X1)) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f116])).
fof(f116, plain, ? [X0] : (? [X1] : (~ ((v1_filter_2(X1, X0) & ~ (k17_filter_2(X0) = X1)) <=> r2_filter_2(X0, X1)) & m2_filter_2(X1, X0)) & (l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => ((v1_filter_2(X1, X0) & ~ (k17_filter_2(X0) = X1)) <=> r2_filter_2(X0, X1)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => ((v1_filter_2(X1, X0) & ~ (k17_filter_2(X0) = X1)) <=> r2_filter_2(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', t58_filter_2)).
fof(f2659, plain, (v3_struct_0(sK2) | (sK3 = k7_filter_2(sK2, sK3))), inference(subsumption_resolution, [], [f2658, f306])).
fof(f306, plain, v10_lattices(sK2), inference(cnf_transformation, [], [f240])).
fof(f2658, plain, (~ v10_lattices(sK2) | v3_struct_0(sK2) | (sK3 = k7_filter_2(sK2, sK3))), inference(subsumption_resolution, [], [f2651, f308])).
fof(f308, plain, l3_lattices(sK2), inference(cnf_transformation, [], [f240])).
fof(f2651, plain, (~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | (sK3 = k7_filter_2(sK2, sK3))), inference(resolution, [], [f1275, f1118])).
fof(f1118, plain, m2_lattice4(sK3, sK2), inference(subsumption_resolution, [], [f1117, f305])).
fof(f1117, plain, (m2_lattice4(sK3, sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1116, f306])).
fof(f1116, plain, (m2_lattice4(sK3, sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1111, f308])).
fof(f1111, plain, (m2_lattice4(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f371, f309])).
fof(f309, plain, m2_filter_2(sK3, sK2), inference(cnf_transformation, [], [f240])).
fof(f371, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f166])).
fof(f166, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', dt_m2_filter_2)).
fof(f1275, plain, ! [X0, X1] : (~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | (k7_filter_2(X0, X1) = X1)), inference(duplicate_literal_removal, [], [f1263])).
fof(f1263, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f351, f372])).
fof(f372, plain, ! [X0, X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f168])).
fof(f168, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f38])).
fof(f38, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_lattice4(X1, X0) => m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', dt_m2_lattice4)).
fof(f351, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | (k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f142])).
fof(f142, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) => (k7_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', d6_filter_2)).
fof(f1553, plain, ((k7_filter_2(sK2, sK3) = k1_filter_0(k1_lattice2(sK2))) | ~ spl32_33), inference(avatar_component_clause, [], [f1551])).
fof(f1551, plain, (spl32_33 <=> (k7_filter_2(sK2, sK3) = k1_filter_0(k1_lattice2(sK2)))), introduced(avatar_definition, [new_symbols(naming, [spl32_33])])).
fof(f1700, plain, ((k1_filter_0(k1_lattice2(sK2)) = k17_filter_2(k1_lattice2(sK2))) | (~ spl32_45 | ~ spl32_46)), inference(forward_demodulation, [], [f1662, f1668])).
fof(f1668, plain, ((u1_struct_0(k1_lattice2(sK2)) = k17_filter_2(k1_lattice2(sK2))) | ~ spl32_46), inference(avatar_component_clause, [], [f1666])).
fof(f1666, plain, (spl32_46 <=> (u1_struct_0(k1_lattice2(sK2)) = k17_filter_2(k1_lattice2(sK2)))), introduced(avatar_definition, [new_symbols(naming, [spl32_46])])).
fof(f1662, plain, ((k1_filter_0(k1_lattice2(sK2)) = u1_struct_0(k1_lattice2(sK2))) | ~ spl32_45), inference(avatar_component_clause, [], [f1660])).
fof(f1660, plain, (spl32_45 <=> (k1_filter_0(k1_lattice2(sK2)) = u1_struct_0(k1_lattice2(sK2)))), introduced(avatar_definition, [new_symbols(naming, [spl32_45])])).
fof(f68847, plain, ((k1_lattice2(sK2) = g3_lattices(k17_filter_2(k1_lattice2(sK2)), u1_lattices(sK2), u2_lattices(sK2))) | (~ spl32_31 | ~ spl32_46 | ~ spl32_2724 | ~ spl32_2725)), inference(forward_demodulation, [], [f68846, f68482])).
fof(f68482, plain, ((u1_lattices(sK2) = u2_lattices(k1_lattice2(sK2))) | ~ spl32_2725), inference(avatar_component_clause, [], [f68480])).
fof(f68480, plain, (spl32_2725 <=> (u1_lattices(sK2) = u2_lattices(k1_lattice2(sK2)))), introduced(avatar_definition, [new_symbols(naming, [spl32_2725])])).
fof(f68846, plain, ((k1_lattice2(sK2) = g3_lattices(k17_filter_2(k1_lattice2(sK2)), u2_lattices(k1_lattice2(sK2)), u2_lattices(sK2))) | (~ spl32_31 | ~ spl32_46 | ~ spl32_2724)), inference(forward_demodulation, [], [f4787, f68473])).
fof(f68473, plain, ((u2_lattices(sK2) = u1_lattices(k1_lattice2(sK2))) | ~ spl32_2724), inference(avatar_component_clause, [], [f68471])).
fof(f68471, plain, (spl32_2724 <=> (u2_lattices(sK2) = u1_lattices(k1_lattice2(sK2)))), introduced(avatar_definition, [new_symbols(naming, [spl32_2724])])).
fof(f4787, plain, ((k1_lattice2(sK2) = g3_lattices(k17_filter_2(k1_lattice2(sK2)), u2_lattices(k1_lattice2(sK2)), u1_lattices(k1_lattice2(sK2)))) | (~ spl32_31 | ~ spl32_46)), inference(backward_demodulation, [], [f3195, f1668])).
fof(f3195, plain, ((k1_lattice2(sK2) = g3_lattices(u1_struct_0(k1_lattice2(sK2)), u2_lattices(k1_lattice2(sK2)), u1_lattices(k1_lattice2(sK2)))) | ~ spl32_31), inference(subsumption_resolution, [], [f3194, f954])).
fof(f954, plain, sP0(sK2), inference(subsumption_resolution, [], [f953, f305])).
fof(f953, plain, (sP0(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f952, f306])).
fof(f952, plain, (sP0(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f943, f307])).
fof(f307, plain, v17_lattices(sK2), inference(cnf_transformation, [], [f240])).
fof(f943, plain, (sP0(sK2) | ~ v17_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f420, f308])).
fof(f420, plain, ! [X0] : (~ l3_lattices(X0) | sP0(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f233])).
fof(f233, plain, ! [X0] : (sP0(X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(definition_folding, [], [f194, e232])).
fof(f232, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(usedef, [], [e232])).
fof(e232, plain, ! [X0] : (sP0(X0) <=> (v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f194, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f193])).
fof(f193, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f113])).
fof(f113, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), inference(pure_predicate_removal, [], [f62])).
fof(f62, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v12_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', fc4_filter_2)).
fof(f3194, plain, ((k1_lattice2(sK2) = g3_lattices(u1_struct_0(k1_lattice2(sK2)), u2_lattices(k1_lattice2(sK2)), u1_lattices(k1_lattice2(sK2)))) | ~ sP0(sK2) | ~ spl32_31), inference(resolution, [], [f1176, f1544])).
fof(f1544, plain, (l3_lattices(k1_lattice2(sK2)) | ~ spl32_31), inference(avatar_component_clause, [], [f1543])).
fof(f1543, plain, (spl32_31 <=> l3_lattices(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_31])])).
fof(f1176, plain, ! [X1] : (~ l3_lattices(k1_lattice2(X1)) | (k1_lattice2(X1) = g3_lattices(u1_struct_0(k1_lattice2(X1)), u2_lattices(k1_lattice2(X1)), u1_lattices(k1_lattice2(X1)))) | ~ sP0(X1)), inference(resolution, [], [f313, f406])).
fof(f406, plain, ! [X0] : (v3_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f263])).
fof(f263, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(nnf_transformation, [], [f232])).
fof(f313, plain, ! [X0] : (~ v3_lattices(X0) | (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ! [X0] : ((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f118])).
fof(f118, plain, ! [X0] : (((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (l3_lattices(X0) => (v3_lattices(X0) => (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', abstractness_v3_lattices)).
fof(f2050, plain, (! [X17, X15, X16] : (~ (k1_lattice2(sK2) = g3_lattices(X15, X16, X17)) | (k17_filter_2(sK2) = X15)) | ~ spl32_109), inference(avatar_component_clause, [], [f2049])).
fof(f2049, plain, (spl32_109 <=> ! [X16, X15, X17] : (~ (k1_lattice2(sK2) = g3_lattices(X15, X16, X17)) | (k17_filter_2(sK2) = X15))), introduced(avatar_definition, [new_symbols(naming, [spl32_109])])).
fof(f69356, plain, (~ spl32_3 | spl32_298), inference(avatar_split_clause, [], [f69355, f4743, f563])).
fof(f563, plain, (spl32_3 <=> r2_filter_2(sK2, sK3)), introduced(avatar_definition, [new_symbols(naming, [spl32_3])])).
fof(f4743, plain, (spl32_298 <=> v1_filter_0(sK3, k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_298])])).
fof(f69355, plain, (v1_filter_0(sK3, k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3)), inference(subsumption_resolution, [], [f69354, f305])).
fof(f69354, plain, (v1_filter_0(sK3, k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f69353, f306])).
fof(f69353, plain, (v1_filter_0(sK3, k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f69352, f308])).
fof(f69352, plain, (v1_filter_0(sK3, k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f4800, f309])).
fof(f4800, plain, (v1_filter_0(sK3, k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f540, f4793])).
fof(f4793, plain, (sK3 = k15_filter_2(sK2, sK3)), inference(forward_demodulation, [], [f1295, f2660])).
fof(f1295, plain, (k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)), inference(subsumption_resolution, [], [f1294, f305])).
fof(f1294, plain, ((k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1293, f306])).
fof(f1293, plain, ((k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1288, f308])).
fof(f1288, plain, ((k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f531, f309])).
fof(f531, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | (k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f211])).
fof(f211, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f210])).
fof(f210, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | (~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f85])).
fof(f85, plain, ! [X0, X1] : ((m2_filter_2(X1, X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k7_filter_2(X0, X1) = k15_filter_2(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', redefinition_k15_filter_2)).
fof(f540, plain, ! [X0, X1] : (v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ r2_filter_2(X0, X1) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f301])).
fof(f301, plain, ! [X0] : (! [X1] : (((r2_filter_2(X0, X1) | ~ v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) & (v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ r2_filter_2(X0, X1))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f220])).
fof(f220, plain, ! [X0] : (! [X1] : ((r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f219])).
fof(f219, plain, ! [X0] : (! [X1] : ((r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f92])).
fof(f92, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', t33_filter_2)).
fof(f68737, plain, (~ spl32_1 | spl32_28 | ~ spl32_29 | ~ spl32_30 | ~ spl32_31 | ~ spl32_45 | ~ spl32_46 | ~ spl32_109 | ~ spl32_298), inference(avatar_contradiction_clause, [], [f68736])).
fof(f68736, plain, ($false | (~ spl32_1 | spl32_28 | ~ spl32_29 | ~ spl32_30 | ~ spl32_31 | ~ spl32_45 | ~ spl32_46 | ~ spl32_109 | ~ spl32_298)), inference(subsumption_resolution, [], [f68630, f4745])).
fof(f4745, plain, (v1_filter_0(sK3, k1_lattice2(sK2)) | ~ spl32_298), inference(avatar_component_clause, [], [f4743])).
fof(f68630, plain, (~ v1_filter_0(sK3, k1_lattice2(sK2)) | (~ spl32_1 | spl32_28 | ~ spl32_29 | ~ spl32_30 | ~ spl32_31 | ~ spl32_45 | ~ spl32_46 | ~ spl32_109)), inference(backward_demodulation, [], [f38879, f68428])).
fof(f68428, plain, ((sK3 = k17_filter_2(k1_lattice2(sK2))) | (~ spl32_1 | ~ spl32_31 | ~ spl32_46 | ~ spl32_109)), inference(trivial_inequality_removal, [], [f68421])).
fof(f68421, plain, (~ (k1_lattice2(sK2) = k1_lattice2(sK2)) | (sK3 = k17_filter_2(k1_lattice2(sK2))) | (~ spl32_1 | ~ spl32_31 | ~ spl32_46 | ~ spl32_109)), inference(superposition, [], [f4858, f4787])).
fof(f4858, plain, (! [X17, X15, X16] : (~ (k1_lattice2(sK2) = g3_lattices(X15, X16, X17)) | (sK3 = X15)) | (~ spl32_1 | ~ spl32_109)), inference(backward_demodulation, [], [f2050, f557])).
fof(f557, plain, ((k17_filter_2(sK2) = sK3) | ~ spl32_1), inference(avatar_component_clause, [], [f555])).
fof(f38879, plain, (~ v1_filter_0(k17_filter_2(k1_lattice2(sK2)), k1_lattice2(sK2)) | (spl32_28 | ~ spl32_29 | ~ spl32_30 | ~ spl32_31 | ~ spl32_45 | ~ spl32_46)), inference(subsumption_resolution, [], [f38878, f1532])).
fof(f1532, plain, (~ v3_struct_0(k1_lattice2(sK2)) | spl32_28), inference(avatar_component_clause, [], [f1531])).
fof(f1531, plain, (spl32_28 <=> v3_struct_0(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_28])])).
fof(f38878, plain, (~ v1_filter_0(k17_filter_2(k1_lattice2(sK2)), k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl32_29 | ~ spl32_30 | ~ spl32_31 | ~ spl32_45 | ~ spl32_46)), inference(subsumption_resolution, [], [f38877, f1536])).
fof(f1536, plain, (v10_lattices(k1_lattice2(sK2)) | ~ spl32_29), inference(avatar_component_clause, [], [f1535])).
fof(f1535, plain, (spl32_29 <=> v10_lattices(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_29])])).
fof(f38877, plain, (~ v1_filter_0(k17_filter_2(k1_lattice2(sK2)), k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl32_30 | ~ spl32_31 | ~ spl32_45 | ~ spl32_46)), inference(subsumption_resolution, [], [f38876, f1540])).
fof(f1540, plain, (v17_lattices(k1_lattice2(sK2)) | ~ spl32_30), inference(avatar_component_clause, [], [f1539])).
fof(f1539, plain, (spl32_30 <=> v17_lattices(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_30])])).
fof(f38876, plain, (~ v1_filter_0(k17_filter_2(k1_lattice2(sK2)), k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl32_31 | ~ spl32_45 | ~ spl32_46)), inference(subsumption_resolution, [], [f38866, f1544])).
fof(f38866, plain, (~ v1_filter_0(k17_filter_2(k1_lattice2(sK2)), k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl32_45 | ~ spl32_46)), inference(superposition, [], [f569, f1700])).
fof(f569, plain, ! [X0] : (~ v1_filter_0(k1_filter_0(X0), X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(subsumption_resolution, [], [f553, f357])).
fof(f357, plain, ! [X0] : (m1_filter_0(k1_filter_0(X0), X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ! [X0] : (m1_filter_0(k1_filter_0(X0), X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f152])).
fof(f152, plain, ! [X0] : (m1_filter_0(k1_filter_0(X0), X0) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_filter_0(k1_filter_0(X0), X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', dt_k1_filter_0)).
fof(f553, plain, ! [X0] : (~ v1_filter_0(k1_filter_0(X0), X0) | ~ m1_filter_0(k1_filter_0(X0), X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(equality_resolution, [], [f547])).
fof(f547, plain, ! [X0, X1] : (~ (k1_filter_0(X0) = X1) | ~ v1_filter_0(X1, X0) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f304])).
fof(f304, plain, ! [X0] : (! [X1] : ((((v2_filter_0(X1, X0) & ~ (k1_filter_0(X0) = X1)) | ~ v1_filter_0(X1, X0)) & (v1_filter_0(X1, X0) | ~ v2_filter_0(X1, X0) | (k1_filter_0(X0) = X1))) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f303])).
fof(f303, plain, ! [X0] : (! [X1] : ((((v2_filter_0(X1, X0) & ~ (k1_filter_0(X0) = X1)) | ~ v1_filter_0(X1, X0)) & (v1_filter_0(X1, X0) | (~ v2_filter_0(X1, X0) | (k1_filter_0(X0) = X1)))) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f227])).
fof(f227, plain, ! [X0] : (! [X1] : (((v2_filter_0(X1, X0) & ~ (k1_filter_0(X0) = X1)) <=> v1_filter_0(X1, X0)) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f226])).
fof(f226, plain, ! [X0] : (! [X1] : (((v2_filter_0(X1, X0) & ~ (k1_filter_0(X0) = X1)) <=> v1_filter_0(X1, X0)) | ~ m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f96])).
fof(f96, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_0(X1, X0) => ((v2_filter_0(X1, X0) & ~ (k1_filter_0(X0) = X1)) <=> v1_filter_0(X1, X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', t58_filter_0)).
fof(f68485, plain, (spl32_2725 | ~ spl32_31 | ~ spl32_46 | ~ spl32_108), inference(avatar_split_clause, [], [f68430, f2043, f1666, f1543, f68480])).
fof(f2043, plain, (spl32_108 <=> ! [X9, X11, X10] : (~ (k1_lattice2(sK2) = g3_lattices(X9, X10, X11)) | (u1_lattices(sK2) = X10))), introduced(avatar_definition, [new_symbols(naming, [spl32_108])])).
fof(f68430, plain, ((u1_lattices(sK2) = u2_lattices(k1_lattice2(sK2))) | (~ spl32_31 | ~ spl32_46 | ~ spl32_108)), inference(trivial_inequality_removal, [], [f68374])).
fof(f68374, plain, (~ (k1_lattice2(sK2) = k1_lattice2(sK2)) | (u1_lattices(sK2) = u2_lattices(k1_lattice2(sK2))) | (~ spl32_31 | ~ spl32_46 | ~ spl32_108)), inference(superposition, [], [f2044, f4787])).
fof(f2044, plain, (! [X10, X11, X9] : (~ (k1_lattice2(sK2) = g3_lattices(X9, X10, X11)) | (u1_lattices(sK2) = X10)) | ~ spl32_108), inference(avatar_component_clause, [], [f2043])).
fof(f68484, plain, (spl32_2724 | ~ spl32_31 | ~ spl32_46 | ~ spl32_107), inference(avatar_split_clause, [], [f68431, f2037, f1666, f1543, f68471])).
fof(f2037, plain, (spl32_107 <=> ! [X3, X5, X4] : (~ (g3_lattices(X3, X4, X5) = k1_lattice2(sK2)) | (u2_lattices(sK2) = X5))), introduced(avatar_definition, [new_symbols(naming, [spl32_107])])).
fof(f68431, plain, ((u2_lattices(sK2) = u1_lattices(k1_lattice2(sK2))) | (~ spl32_31 | ~ spl32_46 | ~ spl32_107)), inference(trivial_inequality_removal, [], [f68373])).
fof(f68373, plain, (~ (k1_lattice2(sK2) = k1_lattice2(sK2)) | (u2_lattices(sK2) = u1_lattices(k1_lattice2(sK2))) | (~ spl32_31 | ~ spl32_46 | ~ spl32_107)), inference(superposition, [], [f2038, f4787])).
fof(f2038, plain, (! [X4, X5, X3] : (~ (g3_lattices(X3, X4, X5) = k1_lattice2(sK2)) | (u2_lattices(sK2) = X5)) | ~ spl32_107), inference(avatar_component_clause, [], [f2037])).
fof(f4835, plain, (spl32_2 | ~ spl32_297), inference(avatar_split_clause, [], [f4834, f4727, f559])).
fof(f559, plain, (spl32_2 <=> v1_filter_2(sK3, sK2)), introduced(avatar_definition, [new_symbols(naming, [spl32_2])])).
fof(f4727, plain, (spl32_297 <=> v2_filter_0(sK3, k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_297])])).
fof(f4834, plain, (~ v2_filter_0(sK3, k1_lattice2(sK2)) | v1_filter_2(sK3, sK2)), inference(subsumption_resolution, [], [f4833, f305])).
fof(f4833, plain, (~ v2_filter_0(sK3, k1_lattice2(sK2)) | v1_filter_2(sK3, sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f4832, f306])).
fof(f4832, plain, (~ v2_filter_0(sK3, k1_lattice2(sK2)) | v1_filter_2(sK3, sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f4831, f308])).
fof(f4831, plain, (~ v2_filter_0(sK3, k1_lattice2(sK2)) | v1_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f4803, f309])).
fof(f4803, plain, (~ v2_filter_0(sK3, k1_lattice2(sK2)) | v1_filter_2(sK3, sK2) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f544, f4793])).
fof(f544, plain, ! [X0, X1] : (~ v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | v1_filter_2(X1, X0) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f302])).
fof(f302, plain, ! [X0] : (! [X1] : (((v1_filter_2(X1, X0) | ~ v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) & (v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ v1_filter_2(X1, X0))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f223])).
fof(f223, plain, ! [X0] : (! [X1] : ((v1_filter_2(X1, X0) <=> v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f222])).
fof(f222, plain, ! [X0] : (! [X1] : ((v1_filter_2(X1, X0) <=> v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f94])).
fof(f94, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (v1_filter_2(X1, X0) <=> v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', t44_filter_2)).
fof(f4830, plain, (spl32_297 | ~ spl32_3 | ~ spl32_29 | ~ spl32_30), inference(avatar_split_clause, [], [f4829, f1539, f1535, f563, f4727])).
fof(f4829, plain, (v2_filter_0(sK3, k1_lattice2(sK2)) | (~ spl32_3 | ~ spl32_29 | ~ spl32_30)), inference(subsumption_resolution, [], [f4828, f1536])).
fof(f4828, plain, (v2_filter_0(sK3, k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | (~ spl32_3 | ~ spl32_30)), inference(subsumption_resolution, [], [f4827, f1540])).
fof(f4827, plain, (v2_filter_0(sK3, k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | ~ spl32_3), inference(subsumption_resolution, [], [f4826, f564])).
fof(f564, plain, (r2_filter_2(sK2, sK3) | ~ spl32_3), inference(avatar_component_clause, [], [f563])).
fof(f4826, plain, (v2_filter_0(sK3, k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2))), inference(subsumption_resolution, [], [f4825, f305])).
fof(f4825, plain, (v2_filter_0(sK3, k1_lattice2(sK2)) | v3_struct_0(sK2) | ~ r2_filter_2(sK2, sK3) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2))), inference(subsumption_resolution, [], [f4824, f306])).
fof(f4824, plain, (v2_filter_0(sK3, k1_lattice2(sK2)) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ r2_filter_2(sK2, sK3) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2))), inference(subsumption_resolution, [], [f4823, f308])).
fof(f4823, plain, (v2_filter_0(sK3, k1_lattice2(sK2)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ r2_filter_2(sK2, sK3) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2))), inference(subsumption_resolution, [], [f4807, f309])).
fof(f4807, plain, (v2_filter_0(sK3, k1_lattice2(sK2)) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ r2_filter_2(sK2, sK3) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2))), inference(superposition, [], [f1372, f4793])).
fof(f1372, plain, ! [X0, X1] : (v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ r2_filter_2(X0, X1) | ~ v17_lattices(k1_lattice2(X0)) | ~ v10_lattices(k1_lattice2(X0))), inference(subsumption_resolution, [], [f1371, f392])).
fof(f392, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ! [X0] : ((v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f183])).
fof(f183, plain, ! [X0] : ((v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f55])).
fof(f55, plain, ! [X0] : ((l3_lattices(X0) & ~ v3_struct_0(X0)) => (v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', fc1_lattice2)).
fof(f1371, plain, ! [X0, X1] : (~ r2_filter_2(X0, X1) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ v17_lattices(k1_lattice2(X0)) | ~ v10_lattices(k1_lattice2(X0)) | v3_struct_0(k1_lattice2(X0))), inference(subsumption_resolution, [], [f1370, f360])).
fof(f360, plain, ! [X0] : (l3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ! [X0] : ((l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ! [X0] : (l3_lattices(X0) => (l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', dt_k1_lattice2)).
fof(f1370, plain, ! [X0, X1] : (~ r2_filter_2(X0, X1) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ l3_lattices(k1_lattice2(X0)) | ~ v17_lattices(k1_lattice2(X0)) | ~ v10_lattices(k1_lattice2(X0)) | v3_struct_0(k1_lattice2(X0))), inference(subsumption_resolution, [], [f1369, f1253])).
fof(f1253, plain, ! [X0, X1] : (m1_filter_0(k15_filter_2(X1, X0), k1_lattice2(X1)) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ m2_filter_2(X0, X1) | ~ v10_lattices(k1_lattice2(X1))), inference(subsumption_resolution, [], [f1252, f392])).
fof(f1252, plain, ! [X0, X1] : (~ m2_filter_2(X0, X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | m1_filter_0(k15_filter_2(X1, X0), k1_lattice2(X1)) | ~ v10_lattices(k1_lattice2(X1)) | v3_struct_0(k1_lattice2(X1))), inference(subsumption_resolution, [], [f1249, f360])).
fof(f1249, plain, ! [X0, X1] : (~ m2_filter_2(X0, X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | m1_filter_0(k15_filter_2(X1, X0), k1_lattice2(X1)) | ~ l3_lattices(k1_lattice2(X1)) | ~ v10_lattices(k1_lattice2(X1)) | v3_struct_0(k1_lattice2(X1))), inference(resolution, [], [f355, f533])).
fof(f533, plain, ! [X0, X1] : (~ m1_filter_2(X1, X0) | m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f299])).
fof(f299, plain, ! [X0] : (! [X1] : ((m1_filter_2(X1, X0) | ~ m1_filter_0(X1, X0)) & (m1_filter_0(X1, X0) | ~ m1_filter_2(X1, X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f215])).
fof(f215, plain, ! [X0] : (! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f214])).
fof(f214, plain, ! [X0] : (! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f87])).
fof(f87, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', redefinition_m1_filter_2)).
fof(f355, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f148])).
fof(f148, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | (~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0, X1] : ((m2_filter_2(X1, X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', dt_k15_filter_2)).
fof(f1369, plain, ! [X0, X1] : (~ r2_filter_2(X0, X1) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ m1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ l3_lattices(k1_lattice2(X0)) | ~ v17_lattices(k1_lattice2(X0)) | ~ v10_lattices(k1_lattice2(X0)) | v3_struct_0(k1_lattice2(X0))), inference(resolution, [], [f540, f548])).
fof(f548, plain, ! [X0, X1] : (~ v1_filter_0(X1, X0) | v2_filter_0(X1, X0) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f304])).
fof(f4781, plain, (spl32_45 | spl32_28 | ~ spl32_29), inference(avatar_split_clause, [], [f2687, f1535, f1531, f1660])).
fof(f2687, plain, ((k1_filter_0(k1_lattice2(sK2)) = u1_struct_0(k1_lattice2(sK2))) | (spl32_28 | ~ spl32_29)), inference(subsumption_resolution, [], [f2686, f308])).
fof(f2686, plain, ((k1_filter_0(k1_lattice2(sK2)) = u1_struct_0(k1_lattice2(sK2))) | ~ l3_lattices(sK2) | (spl32_28 | ~ spl32_29)), inference(subsumption_resolution, [], [f2683, f1532])).
fof(f2683, plain, ((k1_filter_0(k1_lattice2(sK2)) = u1_struct_0(k1_lattice2(sK2))) | v3_struct_0(k1_lattice2(sK2)) | ~ l3_lattices(sK2) | ~ spl32_29), inference(resolution, [], [f973, f1536])).
fof(f973, plain, ! [X0] : (~ v10_lattices(k1_lattice2(X0)) | (u1_struct_0(k1_lattice2(X0)) = k1_filter_0(k1_lattice2(X0))) | v3_struct_0(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(resolution, [], [f349, f360])).
fof(f349, plain, ! [X0] : (~ l3_lattices(X0) | (u1_struct_0(X0) = k1_filter_0(X0)) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ! [X0] : ((u1_struct_0(X0) = k1_filter_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f139])).
fof(f139, plain, ! [X0] : ((u1_struct_0(X0) = k1_filter_0(X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (u1_struct_0(X0) = k1_filter_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', d2_filter_0)).
fof(f4780, plain, (spl32_46 | spl32_28 | ~ spl32_29), inference(avatar_split_clause, [], [f2698, f1535, f1531, f1666])).
fof(f2698, plain, ((u1_struct_0(k1_lattice2(sK2)) = k17_filter_2(k1_lattice2(sK2))) | (spl32_28 | ~ spl32_29)), inference(subsumption_resolution, [], [f2697, f308])).
fof(f2697, plain, ((u1_struct_0(k1_lattice2(sK2)) = k17_filter_2(k1_lattice2(sK2))) | ~ l3_lattices(sK2) | (spl32_28 | ~ spl32_29)), inference(subsumption_resolution, [], [f2694, f1532])).
fof(f2694, plain, ((u1_struct_0(k1_lattice2(sK2)) = k17_filter_2(k1_lattice2(sK2))) | v3_struct_0(k1_lattice2(sK2)) | ~ l3_lattices(sK2) | ~ spl32_29), inference(resolution, [], [f995, f1536])).
fof(f995, plain, ! [X0] : (~ v10_lattices(k1_lattice2(X0)) | (u1_struct_0(k1_lattice2(X0)) = k17_filter_2(k1_lattice2(X0))) | v3_struct_0(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(resolution, [], [f352, f360])).
fof(f352, plain, ! [X0] : (~ l3_lattices(X0) | (k17_filter_2(X0) = u1_struct_0(X0)) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ! [X0] : ((k17_filter_2(X0) = u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f144])).
fof(f144, plain, ! [X0] : ((k17_filter_2(X0) = u1_struct_0(X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k17_filter_2(X0) = u1_struct_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', d8_filter_2)).
fof(f2542, plain, spl32_106, inference(avatar_split_clause, [], [f2541, f2033])).
fof(f2033, plain, (spl32_106 <=> m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_106])])).
fof(f2541, plain, m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)), inference(subsumption_resolution, [], [f2381, f586])).
fof(f586, plain, l2_lattices(sK2), inference(resolution, [], [f365, f308])).
fof(f365, plain, ! [X0] : (~ l3_lattices(X0) | l2_lattices(X0)), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ! [X0] : ((l2_lattices(X0) & l1_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', dt_l3_lattices)).
fof(f2381, plain, (m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ l2_lattices(sK2)), inference(superposition, [], [f1142, f1006])).
fof(f1006, plain, (k17_filter_2(sK2) = u1_struct_0(sK2)), inference(subsumption_resolution, [], [f1005, f305])).
fof(f1005, plain, ((k17_filter_2(sK2) = u1_struct_0(sK2)) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f996, f306])).
fof(f996, plain, ((k17_filter_2(sK2) = u1_struct_0(sK2)) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f352, f308])).
fof(f1142, plain, ! [X1] : (m1_relset_1(u2_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l2_lattices(X1)), inference(resolution, [], [f379, f535])).
fof(f535, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f300])).
fof(f300, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f88])).
fof(f88, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', redefinition_m2_relset_1)).
fof(f379, plain, ! [X0] : (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f172])).
fof(f172, plain, ! [X0] : ((m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0))) | ~ l2_lattices(X0)), inference(ennf_transformation, [], [f42])).
fof(f42, plain, ! [X0] : (l2_lattices(X0) => (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', dt_u2_lattices)).
fof(f2376, plain, spl32_105, inference(avatar_contradiction_clause, [], [f2375])).
fof(f2375, plain, ($false | spl32_105), inference(subsumption_resolution, [], [f2374, f586])).
fof(f2374, plain, (~ l2_lattices(sK2) | spl32_105), inference(resolution, [], [f2031, f377])).
fof(f377, plain, ! [X0] : (v1_funct_1(u2_lattices(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f172])).
fof(f2031, plain, (~ v1_funct_1(u2_lattices(sK2)) | spl32_105), inference(avatar_component_clause, [], [f2029])).
fof(f2029, plain, (spl32_105 <=> v1_funct_1(u2_lattices(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_105])])).
fof(f2373, plain, spl32_104, inference(avatar_split_clause, [], [f2372, f2025])).
fof(f2025, plain, (spl32_104 <=> m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_104])])).
fof(f2372, plain, m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)), inference(subsumption_resolution, [], [f2347, f576])).
fof(f576, plain, l1_lattices(sK2), inference(resolution, [], [f364, f308])).
fof(f364, plain, ! [X0] : (~ l3_lattices(X0) | l1_lattices(X0)), inference(cnf_transformation, [], [f161])).
fof(f2347, plain, (m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ l1_lattices(sK2)), inference(superposition, [], [f1124, f1006])).
fof(f1124, plain, ! [X1] : (m1_relset_1(u1_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l1_lattices(X1)), inference(resolution, [], [f376, f535])).
fof(f376, plain, ! [X0] : (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ! [X0] : ((m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0))) | ~ l1_lattices(X0)), inference(ennf_transformation, [], [f40])).
fof(f40, plain, ! [X0] : (l1_lattices(X0) => (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', dt_u1_lattices)).
fof(f2263, plain, (spl32_32 | spl32_28 | ~ spl32_29 | ~ spl32_31), inference(avatar_split_clause, [], [f2262, f1543, f1535, f1531, f1547])).
fof(f1547, plain, (spl32_32 <=> m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_32])])).
fof(f2262, plain, (m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | (spl32_28 | ~ spl32_29 | ~ spl32_31)), inference(subsumption_resolution, [], [f2261, f1532])).
fof(f2261, plain, (m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl32_29 | ~ spl32_31)), inference(subsumption_resolution, [], [f2260, f1536])).
fof(f2260, plain, (m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | ~ spl32_31), inference(subsumption_resolution, [], [f2246, f1544])).
fof(f2246, plain, (m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))), inference(resolution, [], [f1516, f533])).
fof(f1516, plain, m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)), inference(subsumption_resolution, [], [f1515, f305])).
fof(f1515, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1514, f306])).
fof(f1514, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1513, f308])).
fof(f1513, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1502, f309])).
fof(f1502, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f355, f1295])).
fof(f2054, plain, spl32_103, inference(avatar_contradiction_clause, [], [f2053])).
fof(f2053, plain, ($false | spl32_103), inference(subsumption_resolution, [], [f2052, f576])).
fof(f2052, plain, (~ l1_lattices(sK2) | spl32_103), inference(resolution, [], [f2023, f374])).
fof(f374, plain, ! [X0] : (v1_funct_1(u1_lattices(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f171])).
fof(f2023, plain, (~ v1_funct_1(u1_lattices(sK2)) | spl32_103), inference(avatar_component_clause, [], [f2021])).
fof(f2021, plain, (spl32_103 <=> v1_funct_1(u1_lattices(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl32_103])])).
fof(f2051, plain, (~ spl32_103 | ~ spl32_104 | ~ spl32_105 | ~ spl32_106 | spl32_109), inference(avatar_split_clause, [], [f2047, f2049, f2033, f2029, f2025, f2021])).
fof(f2047, plain, ! [X17, X15, X16] : (~ (k1_lattice2(sK2) = g3_lattices(X15, X16, X17)) | (k17_filter_2(sK2) = X15) | ~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u2_lattices(sK2)) | ~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u1_lattices(sK2))), inference(subsumption_resolution, [], [f2046, f1120])).
fof(f1120, plain, v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)), inference(subsumption_resolution, [], [f1119, f576])).
fof(f1119, plain, (v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ l1_lattices(sK2)), inference(superposition, [], [f375, f1006])).
fof(f375, plain, ! [X0] : (v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f171])).
fof(f2046, plain, ! [X17, X15, X16] : (~ (k1_lattice2(sK2) = g3_lattices(X15, X16, X17)) | (k17_filter_2(sK2) = X15) | ~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u2_lattices(sK2)) | ~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u1_lattices(sK2))), inference(subsumption_resolution, [], [f2017, f1139])).
fof(f1139, plain, v1_funct_2(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)), inference(subsumption_resolution, [], [f1137, f586])).
fof(f1137, plain, (v1_funct_2(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ l2_lattices(sK2)), inference(superposition, [], [f378, f1006])).
fof(f378, plain, ! [X0] : (v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f172])).
fof(f2017, plain, ! [X17, X15, X16] : (~ (k1_lattice2(sK2) = g3_lattices(X15, X16, X17)) | (k17_filter_2(sK2) = X15) | ~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_2(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u2_lattices(sK2)) | ~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u1_lattices(sK2))), inference(superposition, [], [f438, f1096])).
fof(f1096, plain, (k1_lattice2(sK2) = g3_lattices(k17_filter_2(sK2), u1_lattices(sK2), u2_lattices(sK2))), inference(forward_demodulation, [], [f1087, f1006])).
fof(f1087, plain, (k1_lattice2(sK2) = g3_lattices(u1_struct_0(sK2), u1_lattices(sK2), u2_lattices(sK2))), inference(resolution, [], [f350, f308])).
fof(f350, plain, ! [X0] : (~ l3_lattices(X0) | (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ! [X0] : ((k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : (l3_lattices(X0) => (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', d2_lattice2)).
fof(f438, plain, ! [X4, X2, X0, X5, X3, X1] : (~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) | (X0 = X3) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(flattening, [], [f203])).
fof(f203, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | (~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X2) & m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X1)) => ! [X3, X4, X5] : ((g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) => ((X2 = X5) & (X1 = X4) & (X0 = X3)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT322+1.p', free_g3_lattices)).
fof(f2045, plain, (~ spl32_103 | ~ spl32_104 | ~ spl32_105 | ~ spl32_106 | spl32_108), inference(avatar_split_clause, [], [f2041, f2043, f2033, f2029, f2025, f2021])).
fof(f2041, plain, ! [X10, X11, X9] : (~ (k1_lattice2(sK2) = g3_lattices(X9, X10, X11)) | (u1_lattices(sK2) = X10) | ~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u2_lattices(sK2)) | ~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u1_lattices(sK2))), inference(subsumption_resolution, [], [f2040, f1120])).
fof(f2040, plain, ! [X10, X11, X9] : (~ (k1_lattice2(sK2) = g3_lattices(X9, X10, X11)) | (u1_lattices(sK2) = X10) | ~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u2_lattices(sK2)) | ~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u1_lattices(sK2))), inference(subsumption_resolution, [], [f2015, f1139])).
fof(f2015, plain, ! [X10, X11, X9] : (~ (k1_lattice2(sK2) = g3_lattices(X9, X10, X11)) | (u1_lattices(sK2) = X10) | ~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_2(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u2_lattices(sK2)) | ~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u1_lattices(sK2))), inference(superposition, [], [f439, f1096])).
fof(f439, plain, ! [X4, X2, X0, X5, X3, X1] : (~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) | (X1 = X4) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f204])).
fof(f2039, plain, (~ spl32_103 | ~ spl32_104 | ~ spl32_105 | ~ spl32_106 | spl32_107), inference(avatar_split_clause, [], [f2019, f2037, f2033, f2029, f2025, f2021])).
fof(f2019, plain, ! [X4, X5, X3] : (~ (g3_lattices(X3, X4, X5) = k1_lattice2(sK2)) | (u2_lattices(sK2) = X5) | ~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u2_lattices(sK2)) | ~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u1_lattices(sK2))), inference(subsumption_resolution, [], [f2018, f1120])).
fof(f2018, plain, ! [X4, X5, X3] : (~ (g3_lattices(X3, X4, X5) = k1_lattice2(sK2)) | (u2_lattices(sK2) = X5) | ~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u2_lattices(sK2)) | ~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u1_lattices(sK2))), inference(subsumption_resolution, [], [f2013, f1139])).
fof(f2013, plain, ! [X4, X5, X3] : (~ (g3_lattices(X3, X4, X5) = k1_lattice2(sK2)) | (u2_lattices(sK2) = X5) | ~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_2(u2_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u2_lattices(sK2)) | ~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(k17_filter_2(sK2), k17_filter_2(sK2)), k17_filter_2(sK2)) | ~ v1_funct_1(u1_lattices(sK2))), inference(superposition, [], [f440, f1096])).
fof(f440, plain, ! [X4, X2, X0, X5, X3, X1] : (~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) | (X2 = X5) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f204])).
fof(f1698, plain, ~ spl32_28, inference(avatar_contradiction_clause, [], [f1697])).
fof(f1697, plain, ($false | ~ spl32_28), inference(subsumption_resolution, [], [f1691, f954])).
fof(f1691, plain, (~ sP0(sK2) | ~ spl32_28), inference(resolution, [], [f1533, f405])).
fof(f405, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f263])).
fof(f1533, plain, (v3_struct_0(k1_lattice2(sK2)) | ~ spl32_28), inference(avatar_component_clause, [], [f1531])).
fof(f1570, plain, spl32_31, inference(avatar_contradiction_clause, [], [f1569])).
fof(f1569, plain, ($false | spl32_31), inference(subsumption_resolution, [], [f1568, f308])).
fof(f1568, plain, (~ l3_lattices(sK2) | spl32_31), inference(resolution, [], [f1545, f360])).
fof(f1545, plain, (~ l3_lattices(k1_lattice2(sK2)) | spl32_31), inference(avatar_component_clause, [], [f1543])).
fof(f1566, plain, spl32_30, inference(avatar_contradiction_clause, [], [f1565])).
fof(f1565, plain, ($false | spl32_30), inference(subsumption_resolution, [], [f1564, f954])).
fof(f1564, plain, (~ sP0(sK2) | spl32_30), inference(resolution, [], [f1541, f419])).
fof(f419, plain, ! [X0] : (v17_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f263])).
fof(f1541, plain, (~ v17_lattices(k1_lattice2(sK2)) | spl32_30), inference(avatar_component_clause, [], [f1539])).
fof(f1561, plain, spl32_29, inference(avatar_contradiction_clause, [], [f1560])).
fof(f1560, plain, ($false | spl32_29), inference(subsumption_resolution, [], [f1557, f954])).
fof(f1557, plain, (~ sP0(sK2) | spl32_29), inference(resolution, [], [f1537, f413])).
fof(f413, plain, ! [X0] : (v10_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f263])).
fof(f1537, plain, (~ v10_lattices(k1_lattice2(sK2)) | spl32_29), inference(avatar_component_clause, [], [f1535])).
fof(f1554, plain, (spl32_28 | ~ spl32_29 | ~ spl32_30 | ~ spl32_31 | ~ spl32_32 | spl32_33 | ~ spl32_2 | spl32_3), inference(avatar_split_clause, [], [f1529, f563, f559, f1551, f1547, f1543, f1539, f1535, f1531])).
fof(f1529, plain, ((k7_filter_2(sK2, sK3) = k1_filter_0(k1_lattice2(sK2))) | ~ m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl32_2 | spl32_3)), inference(subsumption_resolution, [], [f1528, f1512])).
fof(f1512, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | spl32_3), inference(subsumption_resolution, [], [f1511, f305])).
fof(f1511, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | v3_struct_0(sK2) | spl32_3), inference(subsumption_resolution, [], [f1510, f306])).
fof(f1510, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | spl32_3), inference(subsumption_resolution, [], [f1509, f308])).
fof(f1509, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | spl32_3), inference(subsumption_resolution, [], [f1508, f309])).
fof(f1508, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | spl32_3), inference(subsumption_resolution, [], [f1500, f565])).
fof(f565, plain, (~ r2_filter_2(sK2, sK3) | spl32_3), inference(avatar_component_clause, [], [f563])).
fof(f1500, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | r2_filter_2(sK2, sK3) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f541, f1295])).
fof(f541, plain, ! [X0, X1] : (~ v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | r2_filter_2(X0, X1) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f301])).
fof(f1528, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | (k7_filter_2(sK2, sK3) = k1_filter_0(k1_lattice2(sK2))) | ~ m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | ~ spl32_2), inference(resolution, [], [f1507, f546])).
fof(f546, plain, ! [X0, X1] : (~ v2_filter_0(X1, X0) | v1_filter_0(X1, X0) | (k1_filter_0(X0) = X1) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f304])).
fof(f1507, plain, (v2_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ spl32_2), inference(subsumption_resolution, [], [f1506, f305])).
fof(f1506, plain, (v2_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | v3_struct_0(sK2) | ~ spl32_2), inference(subsumption_resolution, [], [f1505, f306])).
fof(f1505, plain, (v2_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl32_2), inference(subsumption_resolution, [], [f1504, f308])).
fof(f1504, plain, (v2_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl32_2), inference(subsumption_resolution, [], [f1503, f309])).
fof(f1503, plain, (v2_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl32_2), inference(subsumption_resolution, [], [f1499, f560])).
fof(f560, plain, (v1_filter_2(sK3, sK2) | ~ spl32_2), inference(avatar_component_clause, [], [f559])).
fof(f1499, plain, (v2_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ v1_filter_2(sK3, sK2) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f543, f1295])).
fof(f543, plain, ! [X0, X1] : (v2_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ v1_filter_2(X1, X0) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f302])).
fof(f568, plain, (~ spl32_1 | spl32_3), inference(avatar_split_clause, [], [f310, f563, f555])).
fof(f310, plain, (r2_filter_2(sK2, sK3) | ~ (k17_filter_2(sK2) = sK3)), inference(cnf_transformation, [], [f240])).
fof(f567, plain, (spl32_2 | spl32_3), inference(avatar_split_clause, [], [f311, f563, f559])).
fof(f311, plain, (r2_filter_2(sK2, sK3) | v1_filter_2(sK3, sK2)), inference(cnf_transformation, [], [f240])).
fof(f566, plain, (spl32_1 | ~ spl32_2 | ~ spl32_3), inference(avatar_split_clause, [], [f312, f563, f559, f555])).
fof(f312, plain, (~ r2_filter_2(sK2, sK3) | ~ v1_filter_2(sK3, sK2) | (k17_filter_2(sK2) = sK3)), inference(cnf_transformation, [], [f240])).