fof(f317610, plain, $false, inference(avatar_sat_refutation, [], [f570, f575, f580, f584, f585, f1141, f1163, f1185, f1203, f1364, f1735, f1738, f3122, f3128, f3308, f4093, f4096, f4099, f4113, f4117, f7118, f38898, f39216, f39296, f39405, f316601, f316786, f317609])).
fof(f317609, plain, (spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | spl34_180 | spl34_181 | ~ spl34_6172), inference(avatar_contradiction_clause, [], [f317608])).
fof(f317608, plain, ($false | (spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | spl34_180 | spl34_181 | ~ spl34_6172)), inference(subsumption_resolution, [], [f317607, f1127])).
fof(f1127, plain, (~ v3_struct_0(k1_lattice2(sK2)) | spl34_31), inference(avatar_component_clause, [], [f1126])).
fof(f1126, plain, (spl34_31 <=> v3_struct_0(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_31])])).
fof(f317607, plain, (v3_struct_0(k1_lattice2(sK2)) | (~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | spl34_180 | spl34_181 | ~ spl34_6172)), inference(subsumption_resolution, [], [f317606, f1131])).
fof(f1131, plain, (v10_lattices(k1_lattice2(sK2)) | ~ spl34_32), inference(avatar_component_clause, [], [f1130])).
fof(f1130, plain, (spl34_32 <=> v10_lattices(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_32])])).
fof(f317606, plain, (~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl34_33 | ~ spl34_34 | ~ spl34_38 | spl34_180 | spl34_181 | ~ spl34_6172)), inference(subsumption_resolution, [], [f317605, f1248])).
fof(f1248, plain, (v17_lattices(k1_lattice2(sK2)) | ~ spl34_38), inference(avatar_component_clause, [], [f1247])).
fof(f1247, plain, (spl34_38 <=> v17_lattices(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_38])])).
fof(f317605, plain, (~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl34_33 | ~ spl34_34 | spl34_180 | spl34_181 | ~ spl34_6172)), inference(subsumption_resolution, [], [f317604, f1135])).
fof(f1135, plain, (l3_lattices(k1_lattice2(sK2)) | ~ spl34_33), inference(avatar_component_clause, [], [f1134])).
fof(f1134, plain, (spl34_33 <=> l3_lattices(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_33])])).
fof(f317604, plain, (~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl34_34 | spl34_180 | spl34_181 | ~ spl34_6172)), inference(subsumption_resolution, [], [f317603, f2113])).
fof(f2113, plain, (m1_filter_0(sK3, k1_lattice2(sK2)) | ~ spl34_34), inference(backward_demodulation, [], [f1140, f2110])).
fof(f2110, plain, (sK3 = k7_filter_2(sK2, sK3)), inference(subsumption_resolution, [], [f2109, f306])).
fof(f306, plain, ~ v3_struct_0(sK2), inference(cnf_transformation, [], [f239])).
fof(f239, plain, ((((~ r2_hidden(k7_lattices(sK2, sK4), sK3) & ~ r2_hidden(sK4, sK3) & m1_subset_1(sK4, u1_struct_0(sK2))) | (u1_struct_0(sK2) = sK3) | ~ r2_filter_2(sK2, sK3)) & ((! [X3] : (r2_hidden(k7_lattices(sK2, X3), sK3) | r2_hidden(X3, sK3) | ~ m1_subset_1(X3, u1_struct_0(sK2))) & ~ (u1_struct_0(sK2) = sK3)) | r2_filter_2(sK2, sK3)) & m2_filter_2(sK3, sK2)) & l3_lattices(sK2) & v17_lattices(sK2) & v10_lattices(sK2) & ~ v3_struct_0(sK2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2, sK3, sK4])], [f235, f238, f237, f236])).
fof(f236, plain, (? [X0] : (? [X1] : ((? [X2] : (~ r2_hidden(k7_lattices(X0, X2), X1) & ~ r2_hidden(X2, X1) & m1_subset_1(X2, u1_struct_0(X0))) | (u1_struct_0(X0) = X1) | ~ r2_filter_2(X0, X1)) & ((! [X3] : (r2_hidden(k7_lattices(X0, X3), X1) | r2_hidden(X3, X1) | ~ m1_subset_1(X3, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1)) | r2_filter_2(X0, X1)) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (? [X1] : ((? [X2] : (~ r2_hidden(k7_lattices(sK2, X2), X1) & ~ r2_hidden(X2, X1) & m1_subset_1(X2, u1_struct_0(sK2))) | (u1_struct_0(sK2) = X1) | ~ r2_filter_2(sK2, X1)) & ((! [X3] : (r2_hidden(k7_lattices(sK2, X3), X1) | r2_hidden(X3, X1) | ~ m1_subset_1(X3, u1_struct_0(sK2))) & ~ (u1_struct_0(sK2) = X1)) | r2_filter_2(sK2, X1)) & m2_filter_2(X1, sK2)) & l3_lattices(sK2) & v17_lattices(sK2) & v10_lattices(sK2) & ~ v3_struct_0(sK2))), introduced(choice_axiom, [])).
fof(f237, plain, (? [X1] : ((? [X2] : (~ r2_hidden(k7_lattices(sK2, X2), X1) & ~ r2_hidden(X2, X1) & m1_subset_1(X2, u1_struct_0(sK2))) | (u1_struct_0(sK2) = X1) | ~ r2_filter_2(sK2, X1)) & ((! [X3] : (r2_hidden(k7_lattices(sK2, X3), X1) | r2_hidden(X3, X1) | ~ m1_subset_1(X3, u1_struct_0(sK2))) & ~ (u1_struct_0(sK2) = X1)) | r2_filter_2(sK2, X1)) & m2_filter_2(X1, sK2)) => ((? [X2] : (~ r2_hidden(k7_lattices(sK2, X2), sK3) & ~ r2_hidden(X2, sK3) & m1_subset_1(X2, u1_struct_0(sK2))) | (u1_struct_0(sK2) = sK3) | ~ r2_filter_2(sK2, sK3)) & ((! [X3] : (r2_hidden(k7_lattices(sK2, X3), sK3) | r2_hidden(X3, sK3) | ~ m1_subset_1(X3, u1_struct_0(sK2))) & ~ (u1_struct_0(sK2) = sK3)) | r2_filter_2(sK2, sK3)) & m2_filter_2(sK3, sK2))), introduced(choice_axiom, [])).
fof(f238, plain, (? [X2] : (~ r2_hidden(k7_lattices(sK2, X2), sK3) & ~ r2_hidden(X2, sK3) & m1_subset_1(X2, u1_struct_0(sK2))) => (~ r2_hidden(k7_lattices(sK2, sK4), sK3) & ~ r2_hidden(sK4, sK3) & m1_subset_1(sK4, u1_struct_0(sK2)))), introduced(choice_axiom, [])).
fof(f235, plain, ? [X0] : (? [X1] : ((? [X2] : (~ r2_hidden(k7_lattices(X0, X2), X1) & ~ r2_hidden(X2, X1) & m1_subset_1(X2, u1_struct_0(X0))) | (u1_struct_0(X0) = X1) | ~ r2_filter_2(X0, X1)) & ((! [X3] : (r2_hidden(k7_lattices(X0, X3), X1) | r2_hidden(X3, X1) | ~ m1_subset_1(X3, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1)) | r2_filter_2(X0, X1)) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(rectify, [], [f234])).
fof(f234, plain, ? [X0] : (? [X1] : ((? [X2] : (~ r2_hidden(k7_lattices(X0, X2), X1) & ~ r2_hidden(X2, X1) & m1_subset_1(X2, u1_struct_0(X0))) | (u1_struct_0(X0) = X1) | ~ r2_filter_2(X0, X1)) & ((! [X2] : (r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1) | ~ m1_subset_1(X2, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1)) | r2_filter_2(X0, X1)) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f233])).
fof(f233, plain, ? [X0] : (? [X1] : ((((? [X2] : (~ r2_hidden(k7_lattices(X0, X2), X1) & ~ r2_hidden(X2, X1) & m1_subset_1(X2, u1_struct_0(X0))) | (u1_struct_0(X0) = X1)) | ~ r2_filter_2(X0, X1)) & ((! [X2] : (r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1) | ~ m1_subset_1(X2, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1)) | r2_filter_2(X0, X1))) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(nnf_transformation, [], [f116])).
fof(f116, plain, ? [X0] : (? [X1] : (~ (r2_filter_2(X0, X1) <=> (! [X2] : (r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1) | ~ m1_subset_1(X2, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1))) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f115])).
fof(f115, plain, ? [X0] : (? [X1] : (~ (r2_filter_2(X0, X1) <=> (! [X2] : ((r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1)) | ~ m1_subset_1(X2, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1))) & m2_filter_2(X1, X0)) & (l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (r2_filter_2(X0, X1) <=> (! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1))) & ~ (u1_struct_0(X0) = X1))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (r2_filter_2(X0, X1) <=> (! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1))) & ~ (u1_struct_0(X0) = X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', t57_filter_2)).
fof(f2109, plain, (v3_struct_0(sK2) | (sK3 = k7_filter_2(sK2, sK3))), inference(subsumption_resolution, [], [f2108, f307])).
fof(f307, plain, v10_lattices(sK2), inference(cnf_transformation, [], [f239])).
fof(f2108, plain, (~ v10_lattices(sK2) | v3_struct_0(sK2) | (sK3 = k7_filter_2(sK2, sK3))), inference(subsumption_resolution, [], [f2102, f309])).
fof(f309, plain, l3_lattices(sK2), inference(cnf_transformation, [], [f239])).
fof(f2102, plain, (~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | (sK3 = k7_filter_2(sK2, sK3))), inference(resolution, [], [f1109, f1028])).
fof(f1028, plain, m2_lattice4(sK3, sK2), inference(subsumption_resolution, [], [f1027, f306])).
fof(f1027, plain, (m2_lattice4(sK3, sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1026, f307])).
fof(f1026, plain, (m2_lattice4(sK3, sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1023, f309])).
fof(f1023, plain, (m2_lattice4(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f374, f310])).
fof(f310, plain, m2_filter_2(sK3, sK2), inference(cnf_transformation, [], [f239])).
fof(f374, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f165])).
fof(f165, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', dt_m2_filter_2)).
fof(f1109, plain, ! [X0, X1] : (~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | (k7_filter_2(X0, X1) = X1)), inference(duplicate_literal_removal, [], [f1100])).
fof(f1100, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f355, f375])).
fof(f375, plain, ! [X0, X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f168])).
fof(f168, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f167])).
fof(f167, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f38])).
fof(f38, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_lattice4(X1, X0) => m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', dt_m2_lattice4)).
fof(f355, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | (k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f143])).
fof(f143, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) => (k7_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', d6_filter_2)).
fof(f1140, plain, (m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ spl34_34), inference(avatar_component_clause, [], [f1138])).
fof(f1138, plain, (spl34_34 <=> m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_34])])).
fof(f317603, plain, (~ m1_filter_0(sK3, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (spl34_180 | spl34_181 | ~ spl34_6172)), inference(subsumption_resolution, [], [f317602, f3095])).
fof(f3095, plain, (~ (sK3 = u1_struct_0(k1_lattice2(sK2))) | spl34_180), inference(avatar_component_clause, [], [f3094])).
fof(f3094, plain, (spl34_180 <=> (sK3 = u1_struct_0(k1_lattice2(sK2)))), introduced(avatar_definition, [new_symbols(naming, [spl34_180])])).
fof(f317602, plain, ((sK3 = u1_struct_0(k1_lattice2(sK2))) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (spl34_181 | ~ spl34_6172)), inference(subsumption_resolution, [], [f317596, f3104])).
fof(f3104, plain, (~ v1_filter_0(sK3, k1_lattice2(sK2)) | spl34_181), inference(avatar_component_clause, [], [f3103])).
fof(f3103, plain, (spl34_181 <=> v1_filter_0(sK3, k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_181])])).
fof(f317596, plain, (v1_filter_0(sK3, k1_lattice2(sK2)) | (sK3 = u1_struct_0(k1_lattice2(sK2))) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | ~ spl34_6172), inference(resolution, [], [f316785, f551])).
fof(f551, plain, ! [X0, X1] : (~ r2_hidden(sK33(X0, X1), X1) | v1_filter_0(X1, X0) | (u1_struct_0(X0) = X1) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f305])).
fof(f305, plain, ! [X0] : (! [X1] : (((v1_filter_0(X1, X0) | (~ r2_hidden(k7_lattices(X0, sK33(X0, X1)), X1) & ~ r2_hidden(sK33(X0, X1), X1) & m1_subset_1(sK33(X0, X1), u1_struct_0(X0))) | (u1_struct_0(X0) = X1)) & ((! [X3] : (r2_hidden(k7_lattices(X0, X3), X1) | r2_hidden(X3, X1) | ~ m1_subset_1(X3, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1)) | ~ v1_filter_0(X1, X0))) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK33])], [f303, f304])).
fof(f304, plain, ! [X1, X0] : (? [X2] : (~ r2_hidden(k7_lattices(X0, X2), X1) & ~ r2_hidden(X2, X1) & m1_subset_1(X2, u1_struct_0(X0))) => (~ r2_hidden(k7_lattices(X0, sK33(X0, X1)), X1) & ~ r2_hidden(sK33(X0, X1), X1) & m1_subset_1(sK33(X0, X1), u1_struct_0(X0)))), introduced(choice_axiom, [])).
fof(f303, plain, ! [X0] : (! [X1] : (((v1_filter_0(X1, X0) | ? [X2] : (~ r2_hidden(k7_lattices(X0, X2), X1) & ~ r2_hidden(X2, X1) & m1_subset_1(X2, u1_struct_0(X0))) | (u1_struct_0(X0) = X1)) & ((! [X3] : (r2_hidden(k7_lattices(X0, X3), X1) | r2_hidden(X3, X1) | ~ m1_subset_1(X3, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1)) | ~ v1_filter_0(X1, X0))) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(rectify, [], [f302])).
fof(f302, plain, ! [X0] : (! [X1] : (((v1_filter_0(X1, X0) | ? [X2] : (~ r2_hidden(k7_lattices(X0, X2), X1) & ~ r2_hidden(X2, X1) & m1_subset_1(X2, u1_struct_0(X0))) | (u1_struct_0(X0) = X1)) & ((! [X2] : (r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1) | ~ m1_subset_1(X2, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1)) | ~ v1_filter_0(X1, X0))) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f301])).
fof(f301, plain, ! [X0] : (! [X1] : (((v1_filter_0(X1, X0) | (? [X2] : (~ r2_hidden(k7_lattices(X0, X2), X1) & ~ r2_hidden(X2, X1) & m1_subset_1(X2, u1_struct_0(X0))) | (u1_struct_0(X0) = X1))) & ((! [X2] : (r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1) | ~ m1_subset_1(X2, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1)) | ~ v1_filter_0(X1, X0))) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f224])).
fof(f224, plain, ! [X0] : (! [X1] : ((v1_filter_0(X1, X0) <=> (! [X2] : (r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1) | ~ m1_subset_1(X2, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1))) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f223])).
fof(f223, plain, ! [X0] : (! [X1] : ((v1_filter_0(X1, X0) <=> (! [X2] : ((r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1)) | ~ m1_subset_1(X2, u1_struct_0(X0))) & ~ (u1_struct_0(X0) = X1))) | ~ m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f95])).
fof(f95, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_0(X1, X0) => (v1_filter_0(X1, X0) <=> (! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1))) & ~ (u1_struct_0(X0) = X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', t57_filter_0)).
fof(f316785, plain, (r2_hidden(sK33(k1_lattice2(sK2), sK3), sK3) | ~ spl34_6172), inference(avatar_component_clause, [], [f316783])).
fof(f316783, plain, (spl34_6172 <=> r2_hidden(sK33(k1_lattice2(sK2), sK3), sK3)), introduced(avatar_definition, [new_symbols(naming, [spl34_6172])])).
fof(f316786, plain, (spl34_6172 | ~ spl34_3827 | ~ spl34_6 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | spl34_180 | spl34_181 | ~ spl34_541), inference(avatar_split_clause, [], [f255632, f7166, f3103, f3094, f1247, f1138, f1134, f1130, f1126, f582, f201331, f316783])).
fof(f201331, plain, (spl34_3827 <=> m1_subset_1(sK33(k1_lattice2(sK2), sK3), u1_struct_0(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_3827])])).
fof(f582, plain, (spl34_6 <=> ! [X3] : (r2_hidden(k7_lattices(sK2, X3), sK3) | ~ m1_subset_1(X3, u1_struct_0(sK2)) | r2_hidden(X3, sK3))), introduced(avatar_definition, [new_symbols(naming, [spl34_6])])).
fof(f7166, plain, (spl34_541 <=> (k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3)) = k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3)))), introduced(avatar_definition, [new_symbols(naming, [spl34_541])])).
fof(f255632, plain, (~ m1_subset_1(sK33(k1_lattice2(sK2), sK3), u1_struct_0(sK2)) | r2_hidden(sK33(k1_lattice2(sK2), sK3), sK3) | (~ spl34_6 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | spl34_180 | spl34_181 | ~ spl34_541)), inference(resolution, [], [f236582, f583])).
fof(f583, plain, (! [X3] : (r2_hidden(k7_lattices(sK2, X3), sK3) | ~ m1_subset_1(X3, u1_struct_0(sK2)) | r2_hidden(X3, sK3)) | ~ spl34_6), inference(avatar_component_clause, [], [f582])).
fof(f236582, plain, (~ r2_hidden(k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3)), sK3) | (spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | spl34_180 | spl34_181 | ~ spl34_541)), inference(subsumption_resolution, [], [f236581, f1127])).
fof(f236581, plain, (~ r2_hidden(k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3)), sK3) | v3_struct_0(k1_lattice2(sK2)) | (~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | spl34_180 | spl34_181 | ~ spl34_541)), inference(subsumption_resolution, [], [f236580, f1131])).
fof(f236580, plain, (~ r2_hidden(k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3)), sK3) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl34_33 | ~ spl34_34 | ~ spl34_38 | spl34_180 | spl34_181 | ~ spl34_541)), inference(subsumption_resolution, [], [f236579, f1248])).
fof(f236579, plain, (~ r2_hidden(k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3)), sK3) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl34_33 | ~ spl34_34 | spl34_180 | spl34_181 | ~ spl34_541)), inference(subsumption_resolution, [], [f236578, f1135])).
fof(f236578, plain, (~ r2_hidden(k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3)), sK3) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl34_34 | spl34_180 | spl34_181 | ~ spl34_541)), inference(subsumption_resolution, [], [f236577, f2113])).
fof(f236577, plain, (~ r2_hidden(k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3)), sK3) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (spl34_180 | spl34_181 | ~ spl34_541)), inference(subsumption_resolution, [], [f236576, f3095])).
fof(f236576, plain, (~ r2_hidden(k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3)), sK3) | (sK3 = u1_struct_0(k1_lattice2(sK2))) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (spl34_181 | ~ spl34_541)), inference(subsumption_resolution, [], [f236547, f3104])).
fof(f236547, plain, (~ r2_hidden(k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3)), sK3) | v1_filter_0(sK3, k1_lattice2(sK2)) | (sK3 = u1_struct_0(k1_lattice2(sK2))) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | ~ spl34_541), inference(superposition, [], [f552, f7168])).
fof(f7168, plain, ((k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3)) = k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3))) | ~ spl34_541), inference(avatar_component_clause, [], [f7166])).
fof(f552, plain, ! [X0, X1] : (~ r2_hidden(k7_lattices(X0, sK33(X0, X1)), X1) | v1_filter_0(X1, X0) | (u1_struct_0(X0) = X1) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f305])).
fof(f316601, plain, (spl34_2 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | ~ spl34_119 | spl34_181 | spl34_3827), inference(avatar_contradiction_clause, [], [f316600])).
fof(f316600, plain, ($false | (spl34_2 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | ~ spl34_119 | spl34_181 | spl34_3827)), inference(subsumption_resolution, [], [f316599, f2113])).
fof(f316599, plain, (~ m1_filter_0(sK3, k1_lattice2(sK2)) | (spl34_2 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_119 | spl34_181 | spl34_3827)), inference(subsumption_resolution, [], [f316598, f564])).
fof(f564, plain, (~ (u1_struct_0(sK2) = sK3) | spl34_2), inference(avatar_component_clause, [], [f563])).
fof(f563, plain, (spl34_2 <=> (u1_struct_0(sK2) = sK3)), introduced(avatar_definition, [new_symbols(naming, [spl34_2])])).
fof(f316598, plain, ((u1_struct_0(sK2) = sK3) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | (spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_119 | spl34_181 | spl34_3827)), inference(subsumption_resolution, [], [f316477, f3104])).
fof(f316477, plain, (v1_filter_0(sK3, k1_lattice2(sK2)) | (u1_struct_0(sK2) = sK3) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | (spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_119 | spl34_3827)), inference(resolution, [], [f19464, f201333])).
fof(f201333, plain, (~ m1_subset_1(sK33(k1_lattice2(sK2), sK3), u1_struct_0(sK2)) | spl34_3827), inference(avatar_component_clause, [], [f201331])).
fof(f19464, plain, (! [X37] : (m1_subset_1(sK33(k1_lattice2(sK2), X37), u1_struct_0(sK2)) | v1_filter_0(X37, k1_lattice2(sK2)) | (u1_struct_0(sK2) = X37) | ~ m1_filter_0(X37, k1_lattice2(sK2))) | (spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_119)), inference(subsumption_resolution, [], [f19463, f1127])).
fof(f19463, plain, (! [X37] : (m1_subset_1(sK33(k1_lattice2(sK2), X37), u1_struct_0(sK2)) | v1_filter_0(X37, k1_lattice2(sK2)) | (u1_struct_0(sK2) = X37) | ~ m1_filter_0(X37, k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | (~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_119)), inference(subsumption_resolution, [], [f19462, f1131])).
fof(f19462, plain, (! [X37] : (m1_subset_1(sK33(k1_lattice2(sK2), X37), u1_struct_0(sK2)) | v1_filter_0(X37, k1_lattice2(sK2)) | (u1_struct_0(sK2) = X37) | ~ m1_filter_0(X37, k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | (~ spl34_33 | ~ spl34_38 | ~ spl34_119)), inference(subsumption_resolution, [], [f19461, f1248])).
fof(f19461, plain, (! [X37] : (m1_subset_1(sK33(k1_lattice2(sK2), X37), u1_struct_0(sK2)) | v1_filter_0(X37, k1_lattice2(sK2)) | (u1_struct_0(sK2) = X37) | ~ m1_filter_0(X37, k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | (~ spl34_33 | ~ spl34_119)), inference(subsumption_resolution, [], [f19270, f1135])).
fof(f19270, plain, (! [X37] : (m1_subset_1(sK33(k1_lattice2(sK2), X37), u1_struct_0(sK2)) | v1_filter_0(X37, k1_lattice2(sK2)) | (u1_struct_0(sK2) = X37) | ~ m1_filter_0(X37, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | ~ spl34_119), inference(superposition, [], [f550, f8154])).
fof(f8154, plain, ((u1_struct_0(sK2) = u1_struct_0(k1_lattice2(sK2))) | ~ spl34_119), inference(trivial_inequality_removal, [], [f8135])).
fof(f8135, plain, (~ (k1_lattice2(sK2) = k1_lattice2(sK2)) | (u1_struct_0(sK2) = u1_struct_0(k1_lattice2(sK2))) | ~ spl34_119), inference(superposition, [], [f1734, f2130])).
fof(f2130, plain, (k1_lattice2(sK2) = g3_lattices(u1_struct_0(k1_lattice2(sK2)), u2_lattices(k1_lattice2(sK2)), u1_lattices(k1_lattice2(sK2)))), inference(resolution, [], [f1054, f309])).
fof(f1054, plain, ! [X2] : (~ l3_lattices(X2) | (k1_lattice2(X2) = g3_lattices(u1_struct_0(k1_lattice2(X2)), u2_lattices(k1_lattice2(X2)), u1_lattices(k1_lattice2(X2))))), inference(subsumption_resolution, [], [f1046, f360])).
fof(f360, plain, ! [X0] : (l3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ! [X0] : ((l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ! [X0] : (l3_lattices(X0) => (l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', dt_k1_lattice2)).
fof(f1046, plain, ! [X2] : ((k1_lattice2(X2) = g3_lattices(u1_struct_0(k1_lattice2(X2)), u2_lattices(k1_lattice2(X2)), u1_lattices(k1_lattice2(X2)))) | ~ l3_lattices(k1_lattice2(X2)) | ~ l3_lattices(X2)), inference(resolution, [], [f316, f359])).
fof(f359, plain, ! [X0] : (v3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f149])).
fof(f316, plain, ! [X0] : (~ v3_lattices(X0) | (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ! [X0] : ((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f117])).
fof(f117, plain, ! [X0] : (((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (l3_lattices(X0) => (v3_lattices(X0) => (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', abstractness_v3_lattices)).
fof(f1734, plain, (! [X17, X15, X16] : (~ (k1_lattice2(sK2) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK2) = X15)) | ~ spl34_119), inference(avatar_component_clause, [], [f1733])).
fof(f1733, plain, (spl34_119 <=> ! [X16, X15, X17] : (~ (k1_lattice2(sK2) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK2) = X15))), introduced(avatar_definition, [new_symbols(naming, [spl34_119])])).
fof(f550, plain, ! [X0, X1] : (m1_subset_1(sK33(X0, X1), u1_struct_0(X0)) | v1_filter_0(X1, X0) | (u1_struct_0(X0) = X1) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f305])).
fof(f39405, plain, (spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_119 | ~ spl34_539), inference(avatar_contradiction_clause, [], [f39404])).
fof(f39404, plain, ($false | (spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_119 | ~ spl34_539)), inference(subsumption_resolution, [], [f33843, f7141])).
fof(f7141, plain, (! [X3] : ~ m1_subset_1(X3, u1_struct_0(sK2)) | ~ spl34_539), inference(avatar_component_clause, [], [f7140])).
fof(f7140, plain, (spl34_539 <=> ! [X3] : ~ m1_subset_1(X3, u1_struct_0(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_539])])).
fof(f33843, plain, (m1_subset_1(sK12(sK9(k1_lattice2(sK2))), u1_struct_0(sK2)) | (spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_119)), inference(subsumption_resolution, [], [f33842, f1135])).
fof(f33842, plain, (m1_subset_1(sK12(sK9(k1_lattice2(sK2))), u1_struct_0(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | (spl34_31 | ~ spl34_32 | ~ spl34_119)), inference(subsumption_resolution, [], [f33841, f1131])).
fof(f33841, plain, (m1_subset_1(sK12(sK9(k1_lattice2(sK2))), u1_struct_0(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | (spl34_31 | ~ spl34_119)), inference(subsumption_resolution, [], [f33762, f1127])).
fof(f33762, plain, (m1_subset_1(sK12(sK9(k1_lattice2(sK2))), u1_struct_0(sK2)) | v3_struct_0(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ spl34_119), inference(superposition, [], [f11286, f8154])).
fof(f11286, plain, ! [X3] : (m1_subset_1(sK12(sK9(X3)), u1_struct_0(X3)) | v3_struct_0(X3) | ~ v10_lattices(X3) | ~ l3_lattices(X3)), inference(subsumption_resolution, [], [f11283, f984])).
fof(f984, plain, ! [X0] : (~ v1_xboole_0(sK9(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(duplicate_literal_removal, [], [f983])).
fof(f983, plain, ! [X0] : (~ v1_xboole_0(sK9(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f369, f387])).
fof(f387, plain, ! [X0] : (m1_filter_0(sK9(X0), X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f249])).
fof(f249, plain, ! [X0] : (m1_filter_0(sK9(X0), X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9])], [f173, f248])).
fof(f248, plain, ! [X0] : (? [X1] : m1_filter_0(X1, X0) => m1_filter_0(sK9(X0), X0)), introduced(choice_axiom, [])).
fof(f173, plain, ! [X0] : (? [X1] : m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f172])).
fof(f172, plain, ! [X0] : (? [X1] : m1_filter_0(X1, X0) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f47])).
fof(f47, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ? [X1] : m1_filter_0(X1, X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', existence_m1_filter_0)).
fof(f369, plain, ! [X0, X1] : (~ m1_filter_0(X1, X0) | ~ v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ! [X0] : (! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f161])).
fof(f161, plain, ! [X0] : (! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) | ~ m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f33])).
fof(f33, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_0(X1, X0) => (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', dt_m1_filter_0)).
fof(f11283, plain, ! [X3] : (~ v10_lattices(X3) | v3_struct_0(X3) | m1_subset_1(sK12(sK9(X3)), u1_struct_0(X3)) | ~ l3_lattices(X3) | v1_xboole_0(sK9(X3))), inference(resolution, [], [f2375, f921])).
fof(f921, plain, ! [X3] : (r2_hidden(sK12(X3), X3) | v1_xboole_0(X3)), inference(resolution, [], [f541, f390])).
fof(f390, plain, ! [X0] : m1_subset_1(sK12(X0), X0), inference(cnf_transformation, [], [f255])).
fof(f255, plain, ! [X0] : m1_subset_1(sK12(X0), X0), inference(skolemisation, [status(esa), new_symbols(skolem, [sK12])], [f50, f254])).
fof(f254, plain, ! [X0] : (? [X1] : m1_subset_1(X1, X0) => m1_subset_1(sK12(X0), X0)), introduced(choice_axiom, [])).
fof(f50, plain, ! [X0] : ? [X1] : m1_subset_1(X1, X0), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', existence_m1_subset_1)).
fof(f541, plain, ! [X0, X1] : (~ m1_subset_1(X0, X1) | v1_xboole_0(X1) | r2_hidden(X0, X1)), inference(cnf_transformation, [], [f215])).
fof(f215, plain, ! [X0, X1] : (r2_hidden(X0, X1) | v1_xboole_0(X1) | ~ m1_subset_1(X0, X1)), inference(flattening, [], [f214])).
fof(f214, plain, ! [X0, X1] : ((r2_hidden(X0, X1) | v1_xboole_0(X1)) | ~ m1_subset_1(X0, X1)), inference(ennf_transformation, [], [f90])).
fof(f90, plain, ! [X0, X1] : (m1_subset_1(X0, X1) => (r2_hidden(X0, X1) | v1_xboole_0(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', t2_subset)).
fof(f2375, plain, ! [X0, X1] : (~ r2_hidden(X1, sK9(X0)) | ~ v10_lattices(X0) | v3_struct_0(X0) | m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0)), inference(duplicate_literal_removal, [], [f2371])).
fof(f2371, plain, ! [X0, X1] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | m1_subset_1(X1, u1_struct_0(X0)) | ~ r2_hidden(X1, sK9(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f1067, f387])).
fof(f1067, plain, ! [X2, X0, X1] : (~ m1_filter_0(X0, X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | m1_subset_1(X2, u1_struct_0(X1)) | ~ r2_hidden(X2, X0)), inference(resolution, [], [f370, f545])).
fof(f545, plain, ! [X2, X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(X2)) | m1_subset_1(X0, X2) | ~ r2_hidden(X0, X1)), inference(cnf_transformation, [], [f220])).
fof(f220, plain, ! [X0, X1, X2] : (m1_subset_1(X0, X2) | ~ m1_subset_1(X1, k1_zfmisc_1(X2)) | ~ r2_hidden(X0, X1)), inference(flattening, [], [f219])).
fof(f219, plain, ! [X0, X1, X2] : (m1_subset_1(X0, X2) | (~ m1_subset_1(X1, k1_zfmisc_1(X2)) | ~ r2_hidden(X0, X1))), inference(ennf_transformation, [], [f93])).
fof(f93, plain, ! [X0, X1, X2] : ((m1_subset_1(X1, k1_zfmisc_1(X2)) & r2_hidden(X0, X1)) => m1_subset_1(X0, X2)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', t4_subset)).
fof(f370, plain, ! [X0, X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f162])).
fof(f39296, plain, (spl34_181 | spl34_539 | spl34_2 | spl34_541 | ~ spl34_32 | ~ spl34_34 | ~ spl34_38 | ~ spl34_119 | ~ spl34_191), inference(avatar_split_clause, [], [f39295, f3305, f1733, f1247, f1138, f1130, f7166, f563, f7140, f3103])).
fof(f3305, plain, (spl34_191 <=> (sK33(k1_lattice2(sK2), sK3) = k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3)))), introduced(avatar_definition, [new_symbols(naming, [spl34_191])])).
fof(f39295, plain, (! [X3] : ((k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3)) = k7_lattices(sK2, sK33(k1_lattice2(sK2), sK3))) | (u1_struct_0(sK2) = sK3) | ~ m1_subset_1(X3, u1_struct_0(sK2)) | v1_filter_0(sK3, k1_lattice2(sK2))) | (~ spl34_32 | ~ spl34_34 | ~ spl34_38 | ~ spl34_119 | ~ spl34_191)), inference(forward_demodulation, [], [f39294, f3307])).
fof(f3307, plain, ((sK33(k1_lattice2(sK2), sK3) = k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) | ~ spl34_191), inference(avatar_component_clause, [], [f3305])).
fof(f39294, plain, (! [X3] : ((u1_struct_0(sK2) = sK3) | ~ m1_subset_1(X3, u1_struct_0(sK2)) | v1_filter_0(sK3, k1_lattice2(sK2)) | (k7_lattices(sK2, k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) = k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3)))) | (~ spl34_32 | ~ spl34_34 | ~ spl34_38 | ~ spl34_119)), inference(forward_demodulation, [], [f39293, f8154])).
fof(f39293, plain, (! [X3] : (~ m1_subset_1(X3, u1_struct_0(sK2)) | v1_filter_0(sK3, k1_lattice2(sK2)) | (sK3 = u1_struct_0(k1_lattice2(sK2))) | (k7_lattices(sK2, k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) = k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3)))) | (~ spl34_32 | ~ spl34_34 | ~ spl34_38)), inference(subsumption_resolution, [], [f39292, f1131])).
fof(f39292, plain, (! [X3] : (~ m1_subset_1(X3, u1_struct_0(sK2)) | v1_filter_0(sK3, k1_lattice2(sK2)) | (sK3 = u1_struct_0(k1_lattice2(sK2))) | (k7_lattices(sK2, k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) = k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3))) | ~ v10_lattices(k1_lattice2(sK2))) | (~ spl34_34 | ~ spl34_38)), inference(subsumption_resolution, [], [f7131, f1248])).
fof(f7131, plain, (! [X3] : (~ m1_subset_1(X3, u1_struct_0(sK2)) | v1_filter_0(sK3, k1_lattice2(sK2)) | (sK3 = u1_struct_0(k1_lattice2(sK2))) | (k7_lattices(sK2, k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) = k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3))) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2))) | ~ spl34_34), inference(subsumption_resolution, [], [f7130, f306])).
fof(f7130, plain, (! [X3] : (~ m1_subset_1(X3, u1_struct_0(sK2)) | v3_struct_0(sK2) | v1_filter_0(sK3, k1_lattice2(sK2)) | (sK3 = u1_struct_0(k1_lattice2(sK2))) | (k7_lattices(sK2, k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) = k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3))) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2))) | ~ spl34_34), inference(subsumption_resolution, [], [f7129, f307])).
fof(f7129, plain, (! [X3] : (~ m1_subset_1(X3, u1_struct_0(sK2)) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | v1_filter_0(sK3, k1_lattice2(sK2)) | (sK3 = u1_struct_0(k1_lattice2(sK2))) | (k7_lattices(sK2, k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) = k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3))) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2))) | ~ spl34_34), inference(subsumption_resolution, [], [f7128, f308])).
fof(f308, plain, v17_lattices(sK2), inference(cnf_transformation, [], [f239])).
fof(f7128, plain, (! [X3] : (~ m1_subset_1(X3, u1_struct_0(sK2)) | ~ v17_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | v1_filter_0(sK3, k1_lattice2(sK2)) | (sK3 = u1_struct_0(k1_lattice2(sK2))) | (k7_lattices(sK2, k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) = k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3))) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2))) | ~ spl34_34), inference(subsumption_resolution, [], [f3368, f309])).
fof(f3368, plain, (! [X3] : (~ m1_subset_1(X3, u1_struct_0(sK2)) | ~ l3_lattices(sK2) | ~ v17_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | v1_filter_0(sK3, k1_lattice2(sK2)) | (sK3 = u1_struct_0(k1_lattice2(sK2))) | (k7_lattices(sK2, k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) = k7_lattices(k1_lattice2(sK2), sK33(k1_lattice2(sK2), sK3))) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2))) | ~ spl34_34), inference(resolution, [], [f1356, f2113])).
fof(f1356, plain, ! [X12, X13, X11] : (~ m1_filter_0(X12, k1_lattice2(X11)) | ~ m1_subset_1(X13, u1_struct_0(X11)) | ~ l3_lattices(X11) | ~ v17_lattices(X11) | ~ v10_lattices(X11) | v3_struct_0(X11) | v1_filter_0(X12, k1_lattice2(X11)) | (u1_struct_0(k1_lattice2(X11)) = X12) | (k7_lattices(X11, k6_filter_2(X11, sK33(k1_lattice2(X11), X12))) = k7_lattices(k1_lattice2(X11), sK33(k1_lattice2(X11), X12))) | ~ v17_lattices(k1_lattice2(X11)) | ~ v10_lattices(k1_lattice2(X11))), inference(subsumption_resolution, [], [f1355, f395])).
fof(f395, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ! [X0] : ((v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f182])).
fof(f182, plain, ! [X0] : ((v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f55])).
fof(f55, plain, ! [X0] : ((l3_lattices(X0) & ~ v3_struct_0(X0)) => (v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', fc1_lattice2)).
fof(f1355, plain, ! [X12, X13, X11] : ((k7_lattices(X11, k6_filter_2(X11, sK33(k1_lattice2(X11), X12))) = k7_lattices(k1_lattice2(X11), sK33(k1_lattice2(X11), X12))) | ~ m1_subset_1(X13, u1_struct_0(X11)) | ~ l3_lattices(X11) | ~ v17_lattices(X11) | ~ v10_lattices(X11) | v3_struct_0(X11) | v1_filter_0(X12, k1_lattice2(X11)) | (u1_struct_0(k1_lattice2(X11)) = X12) | ~ m1_filter_0(X12, k1_lattice2(X11)) | ~ v17_lattices(k1_lattice2(X11)) | ~ v10_lattices(k1_lattice2(X11)) | v3_struct_0(k1_lattice2(X11))), inference(subsumption_resolution, [], [f1349, f360])).
fof(f1349, plain, ! [X12, X13, X11] : ((k7_lattices(X11, k6_filter_2(X11, sK33(k1_lattice2(X11), X12))) = k7_lattices(k1_lattice2(X11), sK33(k1_lattice2(X11), X12))) | ~ m1_subset_1(X13, u1_struct_0(X11)) | ~ l3_lattices(X11) | ~ v17_lattices(X11) | ~ v10_lattices(X11) | v3_struct_0(X11) | v1_filter_0(X12, k1_lattice2(X11)) | (u1_struct_0(k1_lattice2(X11)) = X12) | ~ m1_filter_0(X12, k1_lattice2(X11)) | ~ l3_lattices(k1_lattice2(X11)) | ~ v17_lattices(k1_lattice2(X11)) | ~ v10_lattices(k1_lattice2(X11)) | v3_struct_0(k1_lattice2(X11))), inference(resolution, [], [f547, f550])).
fof(f547, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, u1_struct_0(k1_lattice2(X0))) | (k7_lattices(X0, k6_filter_2(X0, X2)) = k7_lattices(k1_lattice2(X0), X2)) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f222])).
fof(f222, plain, ! [X0] : (! [X1] : (! [X2] : (((k7_lattices(X0, k6_filter_2(X0, X2)) = k7_lattices(k1_lattice2(X0), X2)) & (k7_lattices(X0, X1) = k7_lattices(k1_lattice2(X0), k5_filter_2(X0, X1)))) | ~ m1_subset_1(X2, u1_struct_0(k1_lattice2(X0)))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f221])).
fof(f221, plain, ! [X0] : (! [X1] : (! [X2] : (((k7_lattices(X0, k6_filter_2(X0, X2)) = k7_lattices(k1_lattice2(X0), X2)) & (k7_lattices(X0, X1) = k7_lattices(k1_lattice2(X0), k5_filter_2(X0, X1)))) | ~ m1_subset_1(X2, u1_struct_0(k1_lattice2(X0)))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f94])).
fof(f94, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2, u1_struct_0(k1_lattice2(X0))) => ((k7_lattices(X0, k6_filter_2(X0, X2)) = k7_lattices(k1_lattice2(X0), X2)) & (k7_lattices(X0, X1) = k7_lattices(k1_lattice2(X0), k5_filter_2(X0, X1))))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', t55_filter_2)).
fof(f39216, plain, (~ spl34_2 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | ~ spl34_119 | ~ spl34_181), inference(avatar_contradiction_clause, [], [f39215])).
fof(f39215, plain, ($false | (~ spl34_2 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | ~ spl34_119 | ~ spl34_181)), inference(subsumption_resolution, [], [f39214, f2113])).
fof(f39214, plain, (~ m1_filter_0(sK3, k1_lattice2(sK2)) | (~ spl34_2 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_119 | ~ spl34_181)), inference(forward_demodulation, [], [f39213, f565])).
fof(f565, plain, ((u1_struct_0(sK2) = sK3) | ~ spl34_2), inference(avatar_component_clause, [], [f563])).
fof(f39213, plain, (~ m1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | (~ spl34_2 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_119 | ~ spl34_181)), inference(subsumption_resolution, [], [f39212, f3105])).
fof(f3105, plain, (v1_filter_0(sK3, k1_lattice2(sK2)) | ~ spl34_181), inference(avatar_component_clause, [], [f3103])).
fof(f39212, plain, (~ v1_filter_0(sK3, k1_lattice2(sK2)) | ~ m1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | (~ spl34_2 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_119)), inference(forward_demodulation, [], [f39211, f565])).
fof(f39211, plain, (~ v1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | ~ m1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | (spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_119)), inference(subsumption_resolution, [], [f39210, f1127])).
fof(f39210, plain, (~ v1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | ~ m1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_119)), inference(subsumption_resolution, [], [f39209, f1131])).
fof(f39209, plain, (~ v1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | ~ m1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl34_33 | ~ spl34_38 | ~ spl34_119)), inference(subsumption_resolution, [], [f39208, f1248])).
fof(f39208, plain, (~ v1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | ~ m1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl34_33 | ~ spl34_119)), inference(subsumption_resolution, [], [f19271, f1135])).
fof(f19271, plain, (~ v1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | ~ m1_filter_0(u1_struct_0(sK2), k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | ~ spl34_119), inference(superposition, [], [f557, f8154])).
fof(f557, plain, ! [X0] : (~ v1_filter_0(u1_struct_0(X0), X0) | ~ m1_filter_0(u1_struct_0(X0), X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(equality_resolution, [], [f548])).
fof(f548, plain, ! [X0, X1] : (~ (u1_struct_0(X0) = X1) | ~ v1_filter_0(X1, X0) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f305])).
fof(f38898, plain, (spl34_4 | spl34_3 | ~ spl34_5 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | ~ spl34_181), inference(avatar_split_clause, [], [f38897, f3103, f1247, f1138, f1134, f1130, f1126, f577, f567, f572])).
fof(f572, plain, (spl34_4 <=> r2_hidden(sK4, sK3)), introduced(avatar_definition, [new_symbols(naming, [spl34_4])])).
fof(f567, plain, (spl34_3 <=> r2_hidden(k7_lattices(sK2, sK4), sK3)), introduced(avatar_definition, [new_symbols(naming, [spl34_3])])).
fof(f577, plain, (spl34_5 <=> m1_subset_1(sK4, u1_struct_0(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_5])])).
fof(f38897, plain, (r2_hidden(sK4, sK3) | (spl34_3 | ~ spl34_5 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_34 | ~ spl34_38 | ~ spl34_181)), inference(subsumption_resolution, [], [f38896, f2113])).
fof(f38896, plain, (r2_hidden(sK4, sK3) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | (spl34_3 | ~ spl34_5 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38 | ~ spl34_181)), inference(subsumption_resolution, [], [f38732, f3105])).
fof(f38732, plain, (r2_hidden(sK4, sK3) | ~ v1_filter_0(sK3, k1_lattice2(sK2)) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | (spl34_3 | ~ spl34_5 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38)), inference(resolution, [], [f7330, f569])).
fof(f569, plain, (~ r2_hidden(k7_lattices(sK2, sK4), sK3) | spl34_3), inference(avatar_component_clause, [], [f567])).
fof(f7330, plain, (! [X2] : (r2_hidden(k7_lattices(sK2, sK4), X2) | r2_hidden(sK4, X2) | ~ v1_filter_0(X2, k1_lattice2(sK2)) | ~ m1_filter_0(X2, k1_lattice2(sK2))) | (~ spl34_5 | spl34_31 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38)), inference(subsumption_resolution, [], [f7329, f1127])).
fof(f7329, plain, (! [X2] : (r2_hidden(k7_lattices(sK2, sK4), X2) | r2_hidden(sK4, X2) | ~ v1_filter_0(X2, k1_lattice2(sK2)) | ~ m1_filter_0(X2, k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | (~ spl34_5 | ~ spl34_32 | ~ spl34_33 | ~ spl34_38)), inference(subsumption_resolution, [], [f7328, f1131])).
fof(f7328, plain, (! [X2] : (r2_hidden(k7_lattices(sK2, sK4), X2) | r2_hidden(sK4, X2) | ~ v1_filter_0(X2, k1_lattice2(sK2)) | ~ m1_filter_0(X2, k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | (~ spl34_5 | ~ spl34_33 | ~ spl34_38)), inference(subsumption_resolution, [], [f7327, f1248])).
fof(f7327, plain, (! [X2] : (r2_hidden(k7_lattices(sK2, sK4), X2) | r2_hidden(sK4, X2) | ~ v1_filter_0(X2, k1_lattice2(sK2)) | ~ m1_filter_0(X2, k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | (~ spl34_5 | ~ spl34_33)), inference(subsumption_resolution, [], [f7326, f1135])).
fof(f7326, plain, (! [X2] : (r2_hidden(k7_lattices(sK2, sK4), X2) | r2_hidden(sK4, X2) | ~ v1_filter_0(X2, k1_lattice2(sK2)) | ~ m1_filter_0(X2, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | ~ spl34_5), inference(subsumption_resolution, [], [f7310, f7224])).
fof(f7224, plain, (m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | ~ spl34_5), inference(subsumption_resolution, [], [f7223, f306])).
fof(f7223, plain, (m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | v3_struct_0(sK2) | ~ spl34_5), inference(subsumption_resolution, [], [f7222, f307])).
fof(f7222, plain, (m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl34_5), inference(subsumption_resolution, [], [f7221, f309])).
fof(f7221, plain, (m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl34_5), inference(subsumption_resolution, [], [f7215, f579])).
fof(f579, plain, (m1_subset_1(sK4, u1_struct_0(sK2)) | ~ spl34_5), inference(avatar_component_clause, [], [f577])).
fof(f7215, plain, (m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | ~ m1_subset_1(sK4, u1_struct_0(sK2)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl34_5), inference(superposition, [], [f361, f7181])).
fof(f7181, plain, ((sK4 = k5_filter_2(sK2, sK4)) | ~ spl34_5), inference(subsumption_resolution, [], [f7180, f306])).
fof(f7180, plain, ((sK4 = k5_filter_2(sK2, sK4)) | v3_struct_0(sK2) | ~ spl34_5), inference(subsumption_resolution, [], [f7179, f307])).
fof(f7179, plain, ((sK4 = k5_filter_2(sK2, sK4)) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl34_5), inference(subsumption_resolution, [], [f7170, f309])).
fof(f7170, plain, ((sK4 = k5_filter_2(sK2, sK4)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl34_5), inference(resolution, [], [f579, f353])).
fof(f353, plain, ! [X0, X1] : (~ m1_subset_1(X1, u1_struct_0(X0)) | (k5_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ! [X0] : (! [X1] : ((k5_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f139])).
fof(f139, plain, ! [X0] : (! [X1] : ((k5_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => (k5_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', d4_filter_2)).
fof(f361, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f150])).
fof(f150, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | (~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ! [X0, X1] : ((m1_subset_1(X1, u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', dt_k5_filter_2)).
fof(f7310, plain, (! [X2] : (r2_hidden(k7_lattices(sK2, sK4), X2) | r2_hidden(sK4, X2) | ~ m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | ~ v1_filter_0(X2, k1_lattice2(sK2)) | ~ m1_filter_0(X2, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | ~ spl34_5), inference(superposition, [], [f549, f7198])).
fof(f7198, plain, ((k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), sK4)) | ~ spl34_5), inference(forward_demodulation, [], [f7197, f7181])).
fof(f7197, plain, ((k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), k5_filter_2(sK2, sK4))) | ~ spl34_5), inference(subsumption_resolution, [], [f7196, f306])).
fof(f7196, plain, ((k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), k5_filter_2(sK2, sK4))) | v3_struct_0(sK2) | ~ spl34_5), inference(subsumption_resolution, [], [f7195, f307])).
fof(f7195, plain, ((k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), k5_filter_2(sK2, sK4))) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl34_5), inference(subsumption_resolution, [], [f7194, f308])).
fof(f7194, plain, ((k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), k5_filter_2(sK2, sK4))) | ~ v17_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl34_5), inference(subsumption_resolution, [], [f7174, f309])).
fof(f7174, plain, ((k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), k5_filter_2(sK2, sK4))) | ~ l3_lattices(sK2) | ~ v17_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl34_5), inference(resolution, [], [f579, f1340])).
fof(f1340, plain, ! [X10, X9] : (~ m1_subset_1(X10, u1_struct_0(X9)) | (k7_lattices(X9, X10) = k7_lattices(k1_lattice2(X9), k5_filter_2(X9, X10))) | ~ l3_lattices(X9) | ~ v17_lattices(X9) | ~ v10_lattices(X9) | v3_struct_0(X9)), inference(resolution, [], [f546, f390])).
fof(f546, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, u1_struct_0(k1_lattice2(X0))) | (k7_lattices(X0, X1) = k7_lattices(k1_lattice2(X0), k5_filter_2(X0, X1))) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f222])).
fof(f549, plain, ! [X0, X3, X1] : (r2_hidden(k7_lattices(X0, X3), X1) | r2_hidden(X3, X1) | ~ m1_subset_1(X3, u1_struct_0(X0)) | ~ v1_filter_0(X1, X0) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f305])).
fof(f7118, plain, (spl34_2 | ~ spl34_119 | ~ spl34_180), inference(avatar_contradiction_clause, [], [f7117])).
fof(f7117, plain, ($false | (spl34_2 | ~ spl34_119 | ~ spl34_180)), inference(subsumption_resolution, [], [f7072, f564])).
fof(f7072, plain, ((u1_struct_0(sK2) = sK3) | (~ spl34_119 | ~ spl34_180)), inference(trivial_inequality_removal, [], [f7071])).
fof(f7071, plain, (~ (k1_lattice2(sK2) = k1_lattice2(sK2)) | (u1_struct_0(sK2) = sK3) | (~ spl34_119 | ~ spl34_180)), inference(superposition, [], [f1734, f3314])).
fof(f3314, plain, ((k1_lattice2(sK2) = g3_lattices(sK3, u2_lattices(k1_lattice2(sK2)), u1_lattices(k1_lattice2(sK2)))) | ~ spl34_180), inference(backward_demodulation, [], [f2130, f3096])).
fof(f3096, plain, ((sK3 = u1_struct_0(k1_lattice2(sK2))) | ~ spl34_180), inference(avatar_component_clause, [], [f3094])).
fof(f4117, plain, spl34_116, inference(avatar_contradiction_clause, [], [f4116])).
fof(f4116, plain, ($false | spl34_116), inference(subsumption_resolution, [], [f4115, f602])).
fof(f602, plain, l2_lattices(sK2), inference(resolution, [], [f368, f309])).
fof(f368, plain, ! [X0] : (~ l3_lattices(X0) | l2_lattices(X0)), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ! [X0] : ((l2_lattices(X0) & l1_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', dt_l3_lattices)).
fof(f4115, plain, (~ l2_lattices(sK2) | spl34_116), inference(resolution, [], [f1723, f1033])).
fof(f1033, plain, ! [X1] : (m1_relset_1(u2_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l2_lattices(X1)), inference(resolution, [], [f382, f537])).
fof(f537, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f299])).
fof(f299, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f87])).
fof(f87, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', redefinition_m2_relset_1)).
fof(f382, plain, ! [X0] : (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ! [X0] : ((m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0))) | ~ l2_lattices(X0)), inference(ennf_transformation, [], [f42])).
fof(f42, plain, ! [X0] : (l2_lattices(X0) => (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', dt_u2_lattices)).
fof(f1723, plain, (~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2)) | spl34_116), inference(avatar_component_clause, [], [f1721])).
fof(f1721, plain, (spl34_116 <=> m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_116])])).
fof(f4113, plain, spl34_115, inference(avatar_contradiction_clause, [], [f4112])).
fof(f4112, plain, ($false | spl34_115), inference(subsumption_resolution, [], [f4111, f602])).
fof(f4111, plain, (~ l2_lattices(sK2) | spl34_115), inference(resolution, [], [f1719, f381])).
fof(f381, plain, ! [X0] : (v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f171])).
fof(f1719, plain, (~ v1_funct_2(u2_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2)) | spl34_115), inference(avatar_component_clause, [], [f1717])).
fof(f1717, plain, (spl34_115 <=> v1_funct_2(u2_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_115])])).
fof(f4099, plain, spl34_114, inference(avatar_contradiction_clause, [], [f4098])).
fof(f4098, plain, ($false | spl34_114), inference(subsumption_resolution, [], [f4097, f602])).
fof(f4097, plain, (~ l2_lattices(sK2) | spl34_114), inference(resolution, [], [f1715, f380])).
fof(f380, plain, ! [X0] : (v1_funct_1(u2_lattices(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f171])).
fof(f1715, plain, (~ v1_funct_1(u2_lattices(sK2)) | spl34_114), inference(avatar_component_clause, [], [f1713])).
fof(f1713, plain, (spl34_114 <=> v1_funct_1(u2_lattices(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_114])])).
fof(f4096, plain, spl34_113, inference(avatar_contradiction_clause, [], [f4095])).
fof(f4095, plain, ($false | spl34_113), inference(subsumption_resolution, [], [f4094, f592])).
fof(f592, plain, l1_lattices(sK2), inference(resolution, [], [f367, f309])).
fof(f367, plain, ! [X0] : (~ l3_lattices(X0) | l1_lattices(X0)), inference(cnf_transformation, [], [f160])).
fof(f4094, plain, (~ l1_lattices(sK2) | spl34_113), inference(resolution, [], [f1711, f1031])).
fof(f1031, plain, ! [X1] : (m1_relset_1(u1_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l1_lattices(X1)), inference(resolution, [], [f379, f537])).
fof(f379, plain, ! [X0] : (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ! [X0] : ((m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0))) | ~ l1_lattices(X0)), inference(ennf_transformation, [], [f40])).
fof(f40, plain, ! [X0] : (l1_lattices(X0) => (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', dt_u1_lattices)).
fof(f1711, plain, (~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2)) | spl34_113), inference(avatar_component_clause, [], [f1709])).
fof(f1709, plain, (spl34_113 <=> m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_113])])).
fof(f4093, plain, spl34_112, inference(avatar_contradiction_clause, [], [f4092])).
fof(f4092, plain, ($false | spl34_112), inference(subsumption_resolution, [], [f4091, f592])).
fof(f4091, plain, (~ l1_lattices(sK2) | spl34_112), inference(resolution, [], [f1707, f378])).
fof(f378, plain, ! [X0] : (v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f170])).
fof(f1707, plain, (~ v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2)) | spl34_112), inference(avatar_component_clause, [], [f1705])).
fof(f1705, plain, (spl34_112 <=> v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_112])])).
fof(f3308, plain, (spl34_191 | spl34_180 | ~ spl34_32 | ~ spl34_34 | ~ spl34_38 | spl34_181), inference(avatar_split_clause, [], [f3303, f3103, f1247, f1138, f1130, f3094, f3305])).
fof(f3303, plain, ((sK3 = u1_struct_0(k1_lattice2(sK2))) | (sK33(k1_lattice2(sK2), sK3) = k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) | (~ spl34_32 | ~ spl34_34 | ~ spl34_38 | spl34_181)), inference(subsumption_resolution, [], [f3302, f306])).
fof(f3302, plain, ((sK3 = u1_struct_0(k1_lattice2(sK2))) | (sK33(k1_lattice2(sK2), sK3) = k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) | v3_struct_0(sK2) | (~ spl34_32 | ~ spl34_34 | ~ spl34_38 | spl34_181)), inference(subsumption_resolution, [], [f3301, f307])).
fof(f3301, plain, ((sK3 = u1_struct_0(k1_lattice2(sK2))) | (sK33(k1_lattice2(sK2), sK3) = k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | (~ spl34_32 | ~ spl34_34 | ~ spl34_38 | spl34_181)), inference(subsumption_resolution, [], [f3300, f309])).
fof(f3300, plain, ((sK3 = u1_struct_0(k1_lattice2(sK2))) | (sK33(k1_lattice2(sK2), sK3) = k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | (~ spl34_32 | ~ spl34_34 | ~ spl34_38 | spl34_181)), inference(subsumption_resolution, [], [f3299, f1131])).
fof(f3299, plain, ((sK3 = u1_struct_0(k1_lattice2(sK2))) | ~ v10_lattices(k1_lattice2(sK2)) | (sK33(k1_lattice2(sK2), sK3) = k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | (~ spl34_34 | ~ spl34_38 | spl34_181)), inference(subsumption_resolution, [], [f3298, f1248])).
fof(f3298, plain, ((sK3 = u1_struct_0(k1_lattice2(sK2))) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | (sK33(k1_lattice2(sK2), sK3) = k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | (~ spl34_34 | spl34_181)), inference(subsumption_resolution, [], [f3292, f3104])).
fof(f3292, plain, ((sK3 = u1_struct_0(k1_lattice2(sK2))) | v1_filter_0(sK3, k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | (sK33(k1_lattice2(sK2), sK3) = k6_filter_2(sK2, sK33(k1_lattice2(sK2), sK3))) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | ~ spl34_34), inference(resolution, [], [f1302, f2113])).
fof(f1302, plain, ! [X2, X3] : (~ m1_filter_0(X2, k1_lattice2(X3)) | (u1_struct_0(k1_lattice2(X3)) = X2) | v1_filter_0(X2, k1_lattice2(X3)) | ~ v17_lattices(k1_lattice2(X3)) | ~ v10_lattices(k1_lattice2(X3)) | (sK33(k1_lattice2(X3), X2) = k6_filter_2(X3, sK33(k1_lattice2(X3), X2))) | ~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3)), inference(subsumption_resolution, [], [f1301, f395])).
fof(f1301, plain, ! [X2, X3] : (v1_filter_0(X2, k1_lattice2(X3)) | (u1_struct_0(k1_lattice2(X3)) = X2) | ~ m1_filter_0(X2, k1_lattice2(X3)) | ~ v17_lattices(k1_lattice2(X3)) | ~ v10_lattices(k1_lattice2(X3)) | v3_struct_0(k1_lattice2(X3)) | (sK33(k1_lattice2(X3), X2) = k6_filter_2(X3, sK33(k1_lattice2(X3), X2))) | ~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3)), inference(subsumption_resolution, [], [f1298, f360])).
fof(f1298, plain, ! [X2, X3] : (v1_filter_0(X2, k1_lattice2(X3)) | (u1_struct_0(k1_lattice2(X3)) = X2) | ~ m1_filter_0(X2, k1_lattice2(X3)) | ~ l3_lattices(k1_lattice2(X3)) | ~ v17_lattices(k1_lattice2(X3)) | ~ v10_lattices(k1_lattice2(X3)) | v3_struct_0(k1_lattice2(X3)) | (sK33(k1_lattice2(X3), X2) = k6_filter_2(X3, sK33(k1_lattice2(X3), X2))) | ~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3)), inference(resolution, [], [f550, f354])).
fof(f354, plain, ! [X0, X1] : (~ m1_subset_1(X1, u1_struct_0(k1_lattice2(X0))) | (k6_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ! [X0] : (! [X1] : ((k6_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, u1_struct_0(k1_lattice2(X0)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f141])).
fof(f141, plain, ! [X0] : (! [X1] : ((k6_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, u1_struct_0(k1_lattice2(X0)))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(k1_lattice2(X0))) => (k6_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', d5_filter_2)).
fof(f3128, plain, (~ spl34_1 | spl34_181), inference(avatar_split_clause, [], [f3127, f3103, f559])).
fof(f559, plain, (spl34_1 <=> r2_filter_2(sK2, sK3)), introduced(avatar_definition, [new_symbols(naming, [spl34_1])])).
fof(f3127, plain, (v1_filter_0(sK3, k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3)), inference(forward_demodulation, [], [f3126, f2110])).
fof(f3126, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3)), inference(subsumption_resolution, [], [f3125, f306])).
fof(f3125, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f3124, f307])).
fof(f3124, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f3123, f309])).
fof(f3123, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1204, f310])).
fof(f1204, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f542, f1116])).
fof(f1116, plain, (k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)), inference(subsumption_resolution, [], [f1115, f306])).
fof(f1115, plain, ((k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1114, f307])).
fof(f1114, plain, ((k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1111, f309])).
fof(f1111, plain, ((k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f534, f310])).
fof(f534, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | (k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f209])).
fof(f209, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | (~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f85])).
fof(f85, plain, ! [X0, X1] : ((m2_filter_2(X1, X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k7_filter_2(X0, X1) = k15_filter_2(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', redefinition_k15_filter_2)).
fof(f542, plain, ! [X0, X1] : (v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ r2_filter_2(X0, X1) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f300])).
fof(f300, plain, ! [X0] : (! [X1] : (((r2_filter_2(X0, X1) | ~ v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) & (v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ r2_filter_2(X0, X1))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f217])).
fof(f217, plain, ! [X0] : (! [X1] : ((r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f216])).
fof(f216, plain, ! [X0] : (! [X1] : ((r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f91])).
fof(f91, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', t33_filter_2)).
fof(f3122, plain, (spl34_1 | ~ spl34_181), inference(avatar_split_clause, [], [f3121, f3103, f559])).
fof(f3121, plain, (~ v1_filter_0(sK3, k1_lattice2(sK2)) | r2_filter_2(sK2, sK3)), inference(forward_demodulation, [], [f3120, f2110])).
fof(f3120, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | r2_filter_2(sK2, sK3)), inference(subsumption_resolution, [], [f3119, f306])).
fof(f3119, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | r2_filter_2(sK2, sK3) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f3118, f307])).
fof(f3118, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | r2_filter_2(sK2, sK3) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f3117, f309])).
fof(f3117, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | r2_filter_2(sK2, sK3) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1206, f310])).
fof(f1206, plain, (~ v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | r2_filter_2(sK2, sK3) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f543, f1116])).
fof(f543, plain, ! [X0, X1] : (~ v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | r2_filter_2(X0, X1) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f300])).
fof(f1738, plain, spl34_111, inference(avatar_contradiction_clause, [], [f1737])).
fof(f1737, plain, ($false | spl34_111), inference(subsumption_resolution, [], [f1736, f592])).
fof(f1736, plain, (~ l1_lattices(sK2) | spl34_111), inference(resolution, [], [f1703, f377])).
fof(f377, plain, ! [X0] : (v1_funct_1(u1_lattices(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f170])).
fof(f1703, plain, (~ v1_funct_1(u1_lattices(sK2)) | spl34_111), inference(avatar_component_clause, [], [f1701])).
fof(f1701, plain, (spl34_111 <=> v1_funct_1(u1_lattices(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl34_111])])).
fof(f1735, plain, (~ spl34_111 | ~ spl34_112 | ~ spl34_113 | ~ spl34_114 | ~ spl34_115 | ~ spl34_116 | spl34_119), inference(avatar_split_clause, [], [f1699, f1733, f1721, f1717, f1713, f1709, f1705, f1701])).
fof(f1699, plain, ! [X17, X15, X16] : (~ (k1_lattice2(sK2) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK2) = X15) | ~ m1_relset_1(u2_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2)) | ~ v1_funct_2(u2_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2)) | ~ v1_funct_1(u2_lattices(sK2)) | ~ m1_relset_1(u1_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2)) | ~ v1_funct_2(u1_lattices(sK2), k2_zfmisc_1(u1_struct_0(sK2), u1_struct_0(sK2)), u1_struct_0(sK2)) | ~ v1_funct_1(u1_lattices(sK2))), inference(superposition, [], [f441, f1012])).
fof(f1012, plain, (k1_lattice2(sK2) = g3_lattices(u1_struct_0(sK2), u1_lattices(sK2), u2_lattices(sK2))), inference(resolution, [], [f352, f309])).
fof(f352, plain, ! [X0] : (~ l3_lattices(X0) | (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ! [X0] : ((k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ! [X0] : (l3_lattices(X0) => (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', d2_lattice2)).
fof(f441, plain, ! [X4, X2, X0, X5, X3, X1] : (~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) | (X0 = X3) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(flattening, [], [f202])).
fof(f202, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | (~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X2) & m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X1)) => ! [X3, X4, X5] : ((g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) => ((X2 = X5) & (X1 = X4) & (X0 = X3)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', free_g3_lattices)).
fof(f1364, plain, spl34_38, inference(avatar_contradiction_clause, [], [f1363])).
fof(f1363, plain, ($false | spl34_38), inference(subsumption_resolution, [], [f1362, f968])).
fof(f968, plain, sP0(sK2), inference(subsumption_resolution, [], [f967, f306])).
fof(f967, plain, (sP0(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f966, f307])).
fof(f966, plain, (sP0(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f957, f308])).
fof(f957, plain, (sP0(sK2) | ~ v17_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f423, f309])).
fof(f423, plain, ! [X0] : (~ l3_lattices(X0) | sP0(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f230])).
fof(f230, plain, ! [X0] : (sP0(X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(definition_folding, [], [f193, e229])).
fof(f229, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(usedef, [], [e229])).
fof(e229, plain, ! [X0] : (sP0(X0) <=> (v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f193, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f192])).
fof(f192, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f112])).
fof(f112, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), inference(pure_predicate_removal, [], [f62])).
fof(f62, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v12_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', fc4_filter_2)).
fof(f1362, plain, (~ sP0(sK2) | spl34_38), inference(resolution, [], [f1249, f422])).
fof(f422, plain, ! [X0] : (v17_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f262])).
fof(f262, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(nnf_transformation, [], [f229])).
fof(f1249, plain, (~ v17_lattices(k1_lattice2(sK2)) | spl34_38), inference(avatar_component_clause, [], [f1247])).
fof(f1203, plain, ~ spl34_31, inference(avatar_contradiction_clause, [], [f1202])).
fof(f1202, plain, ($false | ~ spl34_31), inference(subsumption_resolution, [], [f1196, f968])).
fof(f1196, plain, (~ sP0(sK2) | ~ spl34_31), inference(resolution, [], [f1128, f408])).
fof(f408, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f262])).
fof(f1128, plain, (v3_struct_0(k1_lattice2(sK2)) | ~ spl34_31), inference(avatar_component_clause, [], [f1126])).
fof(f1185, plain, spl34_33, inference(avatar_contradiction_clause, [], [f1184])).
fof(f1184, plain, ($false | spl34_33), inference(subsumption_resolution, [], [f1183, f309])).
fof(f1183, plain, (~ l3_lattices(sK2) | spl34_33), inference(resolution, [], [f1136, f360])).
fof(f1136, plain, (~ l3_lattices(k1_lattice2(sK2)) | spl34_33), inference(avatar_component_clause, [], [f1134])).
fof(f1163, plain, spl34_32, inference(avatar_contradiction_clause, [], [f1162])).
fof(f1162, plain, ($false | spl34_32), inference(subsumption_resolution, [], [f1159, f968])).
fof(f1159, plain, (~ sP0(sK2) | spl34_32), inference(resolution, [], [f1132, f416])).
fof(f416, plain, ! [X0] : (v10_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f262])).
fof(f1132, plain, (~ v10_lattices(k1_lattice2(sK2)) | spl34_32), inference(avatar_component_clause, [], [f1130])).
fof(f1141, plain, (spl34_31 | ~ spl34_32 | ~ spl34_33 | spl34_34), inference(avatar_split_clause, [], [f1122, f1138, f1134, f1130, f1126])).
fof(f1122, plain, (m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))), inference(resolution, [], [f1121, f535])).
fof(f535, plain, ! [X0, X1] : (~ m1_filter_2(X1, X0) | m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f298])).
fof(f298, plain, ! [X0] : (! [X1] : ((m1_filter_2(X1, X0) | ~ m1_filter_0(X1, X0)) & (m1_filter_0(X1, X0) | ~ m1_filter_2(X1, X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f212])).
fof(f212, plain, ! [X0] : (! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f211])).
fof(f211, plain, ! [X0] : (! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f86])).
fof(f86, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', redefinition_m1_filter_2)).
fof(f1121, plain, m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)), inference(subsumption_resolution, [], [f1120, f306])).
fof(f1120, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1119, f307])).
fof(f1119, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1118, f309])).
fof(f1118, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1117, f310])).
fof(f1117, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f358, f1116])).
fof(f358, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f147])).
fof(f147, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | (~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0, X1] : ((m2_filter_2(X1, X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT321+1.p', dt_k15_filter_2)).
fof(f585, plain, (spl34_1 | ~ spl34_2), inference(avatar_split_clause, [], [f311, f563, f559])).
fof(f311, plain, (~ (u1_struct_0(sK2) = sK3) | r2_filter_2(sK2, sK3)), inference(cnf_transformation, [], [f239])).
fof(f584, plain, (spl34_1 | spl34_6), inference(avatar_split_clause, [], [f312, f582, f559])).
fof(f312, plain, ! [X3] : (r2_hidden(k7_lattices(sK2, X3), sK3) | r2_hidden(X3, sK3) | ~ m1_subset_1(X3, u1_struct_0(sK2)) | r2_filter_2(sK2, sK3)), inference(cnf_transformation, [], [f239])).
fof(f580, plain, (~ spl34_1 | spl34_2 | spl34_5), inference(avatar_split_clause, [], [f313, f577, f563, f559])).
fof(f313, plain, (m1_subset_1(sK4, u1_struct_0(sK2)) | (u1_struct_0(sK2) = sK3) | ~ r2_filter_2(sK2, sK3)), inference(cnf_transformation, [], [f239])).
fof(f575, plain, (~ spl34_1 | spl34_2 | ~ spl34_4), inference(avatar_split_clause, [], [f314, f572, f563, f559])).
fof(f314, plain, (~ r2_hidden(sK4, sK3) | (u1_struct_0(sK2) = sK3) | ~ r2_filter_2(sK2, sK3)), inference(cnf_transformation, [], [f239])).
fof(f570, plain, (~ spl34_1 | spl34_2 | ~ spl34_3), inference(avatar_split_clause, [], [f315, f567, f563, f559])).
fof(f315, plain, (~ r2_hidden(k7_lattices(sK2, sK4), sK3) | (u1_struct_0(sK2) = sK3) | ~ r2_filter_2(sK2, sK3)), inference(cnf_transformation, [], [f239])).