fof(f4759, plain, $false, inference(avatar_sat_refutation, [], [f573, f1363, f1972, f2013, f2078, f2118, f2193, f2198, f2222, f3847, f3958, f3969, f4237, f4337, f4748, f4750, f4752, f4754, f4756, f4758])).
fof(f4758, plain, (~ spl37_4 | ~ spl37_206), inference(avatar_contradiction_clause, [], [f4757])).
fof(f4757, plain, ($false | (~ spl37_4 | ~ spl37_206)), inference(subsumption_resolution, [], [f564, f4668])).
fof(f4668, plain, (! [X2] : ~ v7_conlat_1(k15_lattice3(k11_conlat_1(sK3), X2), sK3) | ~ spl37_206), inference(subsumption_resolution, [], [f4667, f329])).
fof(f329, plain, ~ v3_conlat_1(sK3), inference(cnf_transformation, [], [f255])).
fof(f255, plain, (((~ l3_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ v9_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | v7_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ l3_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ v9_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3) | v7_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3)) & m1_subset_1(sK4, k1_zfmisc_1(u1_struct_0(k11_conlat_1(sK3))))) & l2_conlat_1(sK3) & ~ v3_conlat_1(sK3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3, sK4])], [f126, f254, f253])).
fof(f253, plain, (? [X0] : (? [X1] : ((~ l3_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) | ~ v9_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) | v7_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) | ~ l3_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0) | ~ v9_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0) | v7_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0)) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k11_conlat_1(X0))))) & l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (? [X1] : ((~ l3_conlat_1(k15_lattice3(k11_conlat_1(sK3), X1), sK3) | ~ v9_conlat_1(k15_lattice3(k11_conlat_1(sK3), X1), sK3) | v7_conlat_1(k15_lattice3(k11_conlat_1(sK3), X1), sK3) | ~ l3_conlat_1(k16_lattice3(k11_conlat_1(sK3), X1), sK3) | ~ v9_conlat_1(k16_lattice3(k11_conlat_1(sK3), X1), sK3) | v7_conlat_1(k16_lattice3(k11_conlat_1(sK3), X1), sK3)) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k11_conlat_1(sK3))))) & l2_conlat_1(sK3) & ~ v3_conlat_1(sK3))), introduced(choice_axiom, [])).
fof(f254, plain, (? [X1] : ((~ l3_conlat_1(k15_lattice3(k11_conlat_1(sK3), X1), sK3) | ~ v9_conlat_1(k15_lattice3(k11_conlat_1(sK3), X1), sK3) | v7_conlat_1(k15_lattice3(k11_conlat_1(sK3), X1), sK3) | ~ l3_conlat_1(k16_lattice3(k11_conlat_1(sK3), X1), sK3) | ~ v9_conlat_1(k16_lattice3(k11_conlat_1(sK3), X1), sK3) | v7_conlat_1(k16_lattice3(k11_conlat_1(sK3), X1), sK3)) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k11_conlat_1(sK3))))) => ((~ l3_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ v9_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | v7_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ l3_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ v9_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3) | v7_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3)) & m1_subset_1(sK4, k1_zfmisc_1(u1_struct_0(k11_conlat_1(sK3)))))), introduced(choice_axiom, [])).
fof(f126, plain, ? [X0] : (? [X1] : ((~ l3_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) | ~ v9_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) | v7_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) | ~ l3_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0) | ~ v9_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0) | v7_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0)) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k11_conlat_1(X0))))) & l2_conlat_1(X0) & ~ v3_conlat_1(X0)), inference(flattening, [], [f125])).
fof(f125, plain, ? [X0] : (? [X1] : ((~ l3_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) | ~ v9_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) | v7_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) | ~ l3_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0) | ~ v9_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0) | v7_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0)) & m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k11_conlat_1(X0))))) & (l2_conlat_1(X0) & ~ v3_conlat_1(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k11_conlat_1(X0)))) => (l3_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) & v9_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) & ~ v7_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) & l3_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0) & v9_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0) & ~ v7_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k11_conlat_1(X0)))) => (l3_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) & v9_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) & ~ v7_conlat_1(k15_lattice3(k11_conlat_1(X0), X1), X0) & l3_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0) & v9_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0) & ~ v7_conlat_1(k16_lattice3(k11_conlat_1(X0), X1), X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', t4_conlat_2)).
fof(f4667, plain, (! [X2] : (~ v7_conlat_1(k15_lattice3(k11_conlat_1(sK3), X2), sK3) | v3_conlat_1(sK3)) | ~ spl37_206), inference(subsumption_resolution, [], [f4657, f330])).
fof(f330, plain, l2_conlat_1(sK3), inference(cnf_transformation, [], [f255])).
fof(f4657, plain, (! [X2] : (~ v7_conlat_1(k15_lattice3(k11_conlat_1(sK3), X2), sK3) | ~ l2_conlat_1(sK3) | v3_conlat_1(sK3)) | ~ spl37_206), inference(resolution, [], [f4236, f538])).
fof(f538, plain, ! [X0, X1] : (~ r2_hidden(X1, k8_conlat_1(X0)) | ~ v7_conlat_1(X1, X0) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f328])).
fof(f328, plain, ! [X0] : (! [X1] : ((r2_hidden(X1, k8_conlat_1(X0)) | ~ l3_conlat_1(X1, X0) | ~ v9_conlat_1(X1, X0) | v7_conlat_1(X1, X0) | ~ v6_conlat_1(X1, X0)) & ((l3_conlat_1(X1, X0) & v9_conlat_1(X1, X0) & ~ v7_conlat_1(X1, X0) & v6_conlat_1(X1, X0)) | ~ r2_hidden(X1, k8_conlat_1(X0)))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f327])).
fof(f327, plain, ! [X0] : (! [X1] : ((r2_hidden(X1, k8_conlat_1(X0)) | (~ l3_conlat_1(X1, X0) | ~ v9_conlat_1(X1, X0) | v7_conlat_1(X1, X0) | ~ v6_conlat_1(X1, X0))) & ((l3_conlat_1(X1, X0) & v9_conlat_1(X1, X0) & ~ v7_conlat_1(X1, X0) & v6_conlat_1(X1, X0)) | ~ r2_hidden(X1, k8_conlat_1(X0)))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(nnf_transformation, [], [f239])).
fof(f239, plain, ! [X0] : (! [X1] : (r2_hidden(X1, k8_conlat_1(X0)) <=> (l3_conlat_1(X1, X0) & v9_conlat_1(X1, X0) & ~ v7_conlat_1(X1, X0) & v6_conlat_1(X1, X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f238])).
fof(f238, plain, ! [X0] : (! [X1] : (r2_hidden(X1, k8_conlat_1(X0)) <=> (l3_conlat_1(X1, X0) & v9_conlat_1(X1, X0) & ~ v7_conlat_1(X1, X0) & v6_conlat_1(X1, X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f101])).
fof(f101, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => ! [X1] : (r2_hidden(X1, k8_conlat_1(X0)) <=> (l3_conlat_1(X1, X0) & v9_conlat_1(X1, X0) & ~ v7_conlat_1(X1, X0) & v6_conlat_1(X1, X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', t35_conlat_1)).
fof(f4236, plain, (! [X4] : r2_hidden(k15_lattice3(k11_conlat_1(sK3), X4), k8_conlat_1(sK3)) | ~ spl37_206), inference(avatar_component_clause, [], [f4235])).
fof(f4235, plain, (spl37_206 <=> ! [X4] : r2_hidden(k15_lattice3(k11_conlat_1(sK3), X4), k8_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_206])])).
fof(f564, plain, (v7_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ spl37_4), inference(avatar_component_clause, [], [f562])).
fof(f562, plain, (spl37_4 <=> v7_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3)), introduced(avatar_definition, [new_symbols(naming, [spl37_4])])).
fof(f4756, plain, (~ spl37_1 | ~ spl37_206), inference(avatar_contradiction_clause, [], [f4755])).
fof(f4755, plain, ($false | (~ spl37_1 | ~ spl37_206)), inference(subsumption_resolution, [], [f552, f4721])).
fof(f4721, plain, (! [X0] : ~ v7_conlat_1(k16_lattice3(k11_conlat_1(sK3), X0), sK3) | ~ spl37_206), inference(superposition, [], [f4668, f1743])).
fof(f1743, plain, ! [X0] : (k16_lattice3(k11_conlat_1(sK3), X0) = k15_lattice3(k11_conlat_1(sK3), a_2_1_lattice3(k11_conlat_1(sK3), X0))), inference(subsumption_resolution, [], [f1741, f329])).
fof(f1741, plain, ! [X0] : ((k16_lattice3(k11_conlat_1(sK3), X0) = k15_lattice3(k11_conlat_1(sK3), a_2_1_lattice3(k11_conlat_1(sK3), X0))) | v3_conlat_1(sK3)), inference(resolution, [], [f797, f330])).
fof(f797, plain, ! [X0, X1] : (~ l2_conlat_1(X0) | (k16_lattice3(k11_conlat_1(X0), X1) = k15_lattice3(k11_conlat_1(X0), a_2_1_lattice3(k11_conlat_1(X0), X1))) | v3_conlat_1(X0)), inference(subsumption_resolution, [], [f791, f368])).
fof(f368, plain, ! [X0] : (~ v3_struct_0(k11_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f157])).
fof(f157, plain, ! [X0] : ((l3_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f156])).
fof(f156, plain, ! [X0] : ((l3_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (l3_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', dt_k11_conlat_1)).
fof(f791, plain, ! [X0, X1] : ((k16_lattice3(k11_conlat_1(X0), X1) = k15_lattice3(k11_conlat_1(X0), a_2_1_lattice3(k11_conlat_1(X0), X1))) | v3_struct_0(k11_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(resolution, [], [f359, f370])).
fof(f370, plain, ! [X0] : (l3_lattices(k11_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f157])).
fof(f359, plain, ! [X0, X1] : (~ l3_lattices(X0) | (k16_lattice3(X0, X1) = k15_lattice3(X0, a_2_1_lattice3(X0, X1))) | v3_struct_0(X0)), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ! [X0] : (! [X1] : (k16_lattice3(X0, X1) = k15_lattice3(X0, a_2_1_lattice3(X0, X1))) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f146])).
fof(f146, plain, ! [X0] : (! [X1] : (k16_lattice3(X0, X1) = k15_lattice3(X0, a_2_1_lattice3(X0, X1))) | (~ l3_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ! [X0] : ((l3_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (k16_lattice3(X0, X1) = k15_lattice3(X0, a_2_1_lattice3(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', d22_lattice3)).
fof(f552, plain, (v7_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ spl37_1), inference(avatar_component_clause, [], [f550])).
fof(f550, plain, (spl37_1 <=> v7_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3)), introduced(avatar_definition, [new_symbols(naming, [spl37_1])])).
fof(f4754, plain, (spl37_6 | ~ spl37_206), inference(avatar_contradiction_clause, [], [f4753])).
fof(f4753, plain, ($false | (spl37_6 | ~ spl37_206)), inference(subsumption_resolution, [], [f572, f4664])).
fof(f4664, plain, (! [X0] : l3_conlat_1(k15_lattice3(k11_conlat_1(sK3), X0), sK3) | ~ spl37_206), inference(subsumption_resolution, [], [f4663, f329])).
fof(f4663, plain, (! [X0] : (l3_conlat_1(k15_lattice3(k11_conlat_1(sK3), X0), sK3) | v3_conlat_1(sK3)) | ~ spl37_206), inference(subsumption_resolution, [], [f4655, f330])).
fof(f4655, plain, (! [X0] : (l3_conlat_1(k15_lattice3(k11_conlat_1(sK3), X0), sK3) | ~ l2_conlat_1(sK3) | v3_conlat_1(sK3)) | ~ spl37_206), inference(resolution, [], [f4236, f540])).
fof(f540, plain, ! [X0, X1] : (~ r2_hidden(X1, k8_conlat_1(X0)) | l3_conlat_1(X1, X0) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f328])).
fof(f572, plain, (~ l3_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | spl37_6), inference(avatar_component_clause, [], [f570])).
fof(f570, plain, (spl37_6 <=> l3_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3)), introduced(avatar_definition, [new_symbols(naming, [spl37_6])])).
fof(f4752, plain, (spl37_5 | ~ spl37_206), inference(avatar_contradiction_clause, [], [f4751])).
fof(f4751, plain, ($false | (spl37_5 | ~ spl37_206)), inference(subsumption_resolution, [], [f568, f4666])).
fof(f4666, plain, (! [X1] : v9_conlat_1(k15_lattice3(k11_conlat_1(sK3), X1), sK3) | ~ spl37_206), inference(subsumption_resolution, [], [f4665, f329])).
fof(f4665, plain, (! [X1] : (v9_conlat_1(k15_lattice3(k11_conlat_1(sK3), X1), sK3) | v3_conlat_1(sK3)) | ~ spl37_206), inference(subsumption_resolution, [], [f4656, f330])).
fof(f4656, plain, (! [X1] : (v9_conlat_1(k15_lattice3(k11_conlat_1(sK3), X1), sK3) | ~ l2_conlat_1(sK3) | v3_conlat_1(sK3)) | ~ spl37_206), inference(resolution, [], [f4236, f539])).
fof(f539, plain, ! [X0, X1] : (~ r2_hidden(X1, k8_conlat_1(X0)) | v9_conlat_1(X1, X0) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f328])).
fof(f568, plain, (~ v9_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | spl37_5), inference(avatar_component_clause, [], [f566])).
fof(f566, plain, (spl37_5 <=> v9_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3)), introduced(avatar_definition, [new_symbols(naming, [spl37_5])])).
fof(f4750, plain, (spl37_3 | ~ spl37_206), inference(avatar_contradiction_clause, [], [f4749])).
fof(f4749, plain, ($false | (spl37_3 | ~ spl37_206)), inference(subsumption_resolution, [], [f560, f4705])).
fof(f4705, plain, (! [X0] : l3_conlat_1(k16_lattice3(k11_conlat_1(sK3), X0), sK3) | ~ spl37_206), inference(superposition, [], [f4664, f1743])).
fof(f560, plain, (~ l3_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3) | spl37_3), inference(avatar_component_clause, [], [f558])).
fof(f558, plain, (spl37_3 <=> l3_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3)), introduced(avatar_definition, [new_symbols(naming, [spl37_3])])).
fof(f4748, plain, (spl37_2 | ~ spl37_206), inference(avatar_contradiction_clause, [], [f4747])).
fof(f4747, plain, ($false | (spl37_2 | ~ spl37_206)), inference(subsumption_resolution, [], [f556, f4712])).
fof(f4712, plain, (! [X0] : v9_conlat_1(k16_lattice3(k11_conlat_1(sK3), X0), sK3) | ~ spl37_206), inference(superposition, [], [f4666, f1743])).
fof(f556, plain, (~ v9_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3) | spl37_2), inference(avatar_component_clause, [], [f554])).
fof(f554, plain, (spl37_2 <=> v9_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3)), introduced(avatar_definition, [new_symbols(naming, [spl37_2])])).
fof(f4337, plain, ~ spl37_201, inference(avatar_contradiction_clause, [], [f4336])).
fof(f4336, plain, ($false | ~ spl37_201), inference(subsumption_resolution, [], [f4328, f606])).
fof(f606, plain, sP0(sK3), inference(subsumption_resolution, [], [f604, f329])).
fof(f604, plain, (sP0(sK3) | v3_conlat_1(sK3)), inference(resolution, [], [f417, f330])).
fof(f417, plain, ! [X0] : (~ l2_conlat_1(X0) | sP0(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f248])).
fof(f248, plain, ! [X0] : (sP0(X0) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(definition_folding, [], [f187, e247])).
fof(f247, plain, ! [X0] : ((v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ sP0(X0)), inference(usedef, [], [e247])).
fof(e247, plain, ! [X0] : (sP0(X0) <=> (v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f187, plain, ! [X0] : ((v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f186])).
fof(f186, plain, ! [X0] : ((v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f114])).
fof(f114, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0)))), inference(pure_predicate_removal, [], [f59])).
fof(f59, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (v4_lattice3(k11_conlat_1(X0)) & v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', fc1_conlat_2)).
fof(f4328, plain, (~ sP0(sK3) | ~ spl37_201), inference(resolution, [], [f4195, f405])).
fof(f405, plain, ! [X0] : (~ v3_struct_0(k11_conlat_1(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ! [X0] : ((v15_lattices(k11_conlat_1(X0)) & v14_lattices(k11_conlat_1(X0)) & v13_lattices(k11_conlat_1(X0)) & v10_lattices(k11_conlat_1(X0)) & v9_lattices(k11_conlat_1(X0)) & v8_lattices(k11_conlat_1(X0)) & v7_lattices(k11_conlat_1(X0)) & v6_lattices(k11_conlat_1(X0)) & v5_lattices(k11_conlat_1(X0)) & v4_lattices(k11_conlat_1(X0)) & v3_lattices(k11_conlat_1(X0)) & ~ v3_struct_0(k11_conlat_1(X0))) | ~ sP0(X0)), inference(nnf_transformation, [], [f247])).
fof(f4195, plain, (v3_struct_0(k11_conlat_1(sK3)) | ~ spl37_201), inference(avatar_component_clause, [], [f4193])).
fof(f4193, plain, (spl37_201 <=> v3_struct_0(k11_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_201])])).
fof(f4237, plain, (spl37_201 | spl37_206 | ~ spl37_63 | spl37_68 | ~ spl37_169), inference(avatar_split_clause, [], [f4233, f3844, f2016, f1357, f4235, f4193])).
fof(f1357, plain, (spl37_63 <=> ! [X16, X15, X17] : (~ (k11_conlat_1(sK3) = g3_lattices(X15, X16, X17)) | (k8_conlat_1(sK3) = X15))), introduced(avatar_definition, [new_symbols(naming, [spl37_63])])).
fof(f2016, plain, (spl37_68 <=> v1_xboole_0(k8_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_68])])).
fof(f3844, plain, (spl37_169 <=> l3_lattices(k11_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_169])])).
fof(f4233, plain, (! [X4] : (r2_hidden(k15_lattice3(k11_conlat_1(sK3), X4), k8_conlat_1(sK3)) | v3_struct_0(k11_conlat_1(sK3))) | (~ spl37_63 | spl37_68 | ~ spl37_169)), inference(subsumption_resolution, [], [f4232, f3846])).
fof(f3846, plain, (l3_lattices(k11_conlat_1(sK3)) | ~ spl37_169), inference(avatar_component_clause, [], [f3844])).
fof(f4232, plain, (! [X4] : (r2_hidden(k15_lattice3(k11_conlat_1(sK3), X4), k8_conlat_1(sK3)) | v3_struct_0(k11_conlat_1(sK3)) | ~ l3_lattices(k11_conlat_1(sK3))) | (~ spl37_63 | spl37_68)), inference(subsumption_resolution, [], [f4166, f2017])).
fof(f2017, plain, (~ v1_xboole_0(k8_conlat_1(sK3)) | spl37_68), inference(avatar_component_clause, [], [f2016])).
fof(f4166, plain, (! [X4] : (r2_hidden(k15_lattice3(k11_conlat_1(sK3), X4), k8_conlat_1(sK3)) | v3_struct_0(k11_conlat_1(sK3)) | v1_xboole_0(k8_conlat_1(sK3)) | ~ l3_lattices(k11_conlat_1(sK3))) | ~ spl37_63), inference(superposition, [], [f708, f3832])).
fof(f3832, plain, ((u1_struct_0(k11_conlat_1(sK3)) = k8_conlat_1(sK3)) | ~ spl37_63), inference(trivial_inequality_removal, [], [f3831])).
fof(f3831, plain, (~ (k11_conlat_1(sK3) = k11_conlat_1(sK3)) | (u1_struct_0(k11_conlat_1(sK3)) = k8_conlat_1(sK3)) | ~ spl37_63), inference(superposition, [], [f1358, f1765])).
fof(f1765, plain, (k11_conlat_1(sK3) = g3_lattices(u1_struct_0(k11_conlat_1(sK3)), u2_lattices(k11_conlat_1(sK3)), u1_lattices(k11_conlat_1(sK3)))), inference(subsumption_resolution, [], [f1763, f329])).
fof(f1763, plain, ((k11_conlat_1(sK3) = g3_lattices(u1_struct_0(k11_conlat_1(sK3)), u2_lattices(k11_conlat_1(sK3)), u1_lattices(k11_conlat_1(sK3)))) | v3_conlat_1(sK3)), inference(resolution, [], [f784, f330])).
fof(f784, plain, ! [X0] : (~ l2_conlat_1(X0) | (k11_conlat_1(X0) = g3_lattices(u1_struct_0(k11_conlat_1(X0)), u2_lattices(k11_conlat_1(X0)), u1_lattices(k11_conlat_1(X0)))) | v3_conlat_1(X0)), inference(subsumption_resolution, [], [f776, f370])).
fof(f776, plain, ! [X0] : ((k11_conlat_1(X0) = g3_lattices(u1_struct_0(k11_conlat_1(X0)), u2_lattices(k11_conlat_1(X0)), u1_lattices(k11_conlat_1(X0)))) | ~ l3_lattices(k11_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(resolution, [], [f333, f369])).
fof(f369, plain, ! [X0] : (v3_lattices(k11_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f157])).
fof(f333, plain, ! [X0] : (~ v3_lattices(X0) | (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ! [X0] : ((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f127])).
fof(f127, plain, ! [X0] : (((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (l3_lattices(X0) => (v3_lattices(X0) => (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', abstractness_v3_lattices)).
fof(f1358, plain, (! [X17, X15, X16] : (~ (k11_conlat_1(sK3) = g3_lattices(X15, X16, X17)) | (k8_conlat_1(sK3) = X15)) | ~ spl37_63), inference(avatar_component_clause, [], [f1357])).
fof(f708, plain, ! [X0, X1] : (r2_hidden(k15_lattice3(X0, X1), u1_struct_0(X0)) | v3_struct_0(X0) | v1_xboole_0(u1_struct_0(X0)) | ~ l3_lattices(X0)), inference(resolution, [], [f371, f534])).
fof(f534, plain, ! [X0, X1] : (~ m1_subset_1(X0, X1) | v1_xboole_0(X1) | r2_hidden(X0, X1)), inference(cnf_transformation, [], [f236])).
fof(f236, plain, ! [X0, X1] : (r2_hidden(X0, X1) | v1_xboole_0(X1) | ~ m1_subset_1(X0, X1)), inference(flattening, [], [f235])).
fof(f235, plain, ! [X0, X1] : ((r2_hidden(X0, X1) | v1_xboole_0(X1)) | ~ m1_subset_1(X0, X1)), inference(ennf_transformation, [], [f99])).
fof(f99, plain, ! [X0, X1] : (m1_subset_1(X0, X1) => (r2_hidden(X0, X1) | v1_xboole_0(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', t2_subset)).
fof(f371, plain, ! [X0, X1] : (m1_subset_1(k15_lattice3(X0, X1), u1_struct_0(X0)) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ! [X0, X1] : (m1_subset_1(k15_lattice3(X0, X1), u1_struct_0(X0)) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f158])).
fof(f158, plain, ! [X0, X1] : (m1_subset_1(k15_lattice3(X0, X1), u1_struct_0(X0)) | (~ l3_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ! [X0, X1] : ((l3_lattices(X0) & ~ v3_struct_0(X0)) => m1_subset_1(k15_lattice3(X0, X1), u1_struct_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', dt_k15_lattice3)).
fof(f3969, plain, spl37_168, inference(avatar_contradiction_clause, [], [f3968])).
fof(f3968, plain, ($false | spl37_168), inference(subsumption_resolution, [], [f3967, f330])).
fof(f3967, plain, (~ l2_conlat_1(sK3) | spl37_168), inference(subsumption_resolution, [], [f3966, f329])).
fof(f3966, plain, (v3_conlat_1(sK3) | ~ l2_conlat_1(sK3) | spl37_168), inference(resolution, [], [f3842, f623])).
fof(f623, plain, ! [X1] : (l1_lattices(k11_conlat_1(X1)) | v3_conlat_1(X1) | ~ l2_conlat_1(X1)), inference(resolution, [], [f370, f381])).
fof(f381, plain, ! [X0] : (~ l3_lattices(X0) | l1_lattices(X0)), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ! [X0] : ((l2_lattices(X0) & l1_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', dt_l3_lattices)).
fof(f3842, plain, (~ l1_lattices(k11_conlat_1(sK3)) | spl37_168), inference(avatar_component_clause, [], [f3840])).
fof(f3840, plain, (spl37_168 <=> l1_lattices(k11_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_168])])).
fof(f3958, plain, spl37_167, inference(avatar_contradiction_clause, [], [f3957])).
fof(f3957, plain, ($false | spl37_167), inference(subsumption_resolution, [], [f3956, f330])).
fof(f3956, plain, (~ l2_conlat_1(sK3) | spl37_167), inference(subsumption_resolution, [], [f3955, f329])).
fof(f3955, plain, (v3_conlat_1(sK3) | ~ l2_conlat_1(sK3) | spl37_167), inference(resolution, [], [f3838, f622])).
fof(f622, plain, ! [X0] : (l2_lattices(k11_conlat_1(X0)) | v3_conlat_1(X0) | ~ l2_conlat_1(X0)), inference(resolution, [], [f370, f382])).
fof(f382, plain, ! [X0] : (~ l3_lattices(X0) | l2_lattices(X0)), inference(cnf_transformation, [], [f171])).
fof(f3838, plain, (~ l2_lattices(k11_conlat_1(sK3)) | spl37_167), inference(avatar_component_clause, [], [f3836])).
fof(f3836, plain, (spl37_167 <=> l2_lattices(k11_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_167])])).
fof(f3847, plain, (~ spl37_167 | ~ spl37_168 | spl37_169), inference(avatar_split_clause, [], [f3809, f3844, f3840, f3836])).
fof(f3809, plain, (l3_lattices(k11_conlat_1(sK3)) | ~ l1_lattices(k11_conlat_1(sK3)) | ~ l2_lattices(k11_conlat_1(sK3))), inference(superposition, [], [f2301, f1765])).
fof(f2301, plain, ! [X0] : (l3_lattices(g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0))) | ~ l1_lattices(X0) | ~ l2_lattices(X0)), inference(subsumption_resolution, [], [f2300, f388])).
fof(f388, plain, ! [X0] : (v1_funct_1(u2_lattices(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ! [X0] : ((m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0))) | ~ l2_lattices(X0)), inference(ennf_transformation, [], [f44])).
fof(f44, plain, ! [X0] : (l2_lattices(X0) => (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', dt_u2_lattices)).
fof(f2300, plain, ! [X0] : (l3_lattices(g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0))) | ~ v1_funct_1(u2_lattices(X0)) | ~ l1_lattices(X0) | ~ l2_lattices(X0)), inference(subsumption_resolution, [], [f2296, f760])).
fof(f760, plain, ! [X0] : (m1_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(resolution, [], [f390, f530])).
fof(f530, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f323])).
fof(f323, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f96])).
fof(f96, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', redefinition_m2_relset_1)).
fof(f390, plain, ! [X0] : (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f176])).
fof(f2296, plain, ! [X0] : (~ m1_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | l3_lattices(g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0))) | ~ v1_funct_1(u2_lattices(X0)) | ~ l1_lattices(X0) | ~ l2_lattices(X0)), inference(resolution, [], [f1053, f389])).
fof(f389, plain, ! [X0] : (v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f176])).
fof(f1053, plain, ! [X2, X3] : (~ v1_funct_2(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | ~ m1_relset_1(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | l3_lattices(g3_lattices(u1_struct_0(X2), X3, u1_lattices(X2))) | ~ v1_funct_1(X3) | ~ l1_lattices(X2)), inference(subsumption_resolution, [], [f1052, f385])).
fof(f385, plain, ! [X0] : (v1_funct_1(u1_lattices(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ! [X0] : ((m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0))) | ~ l1_lattices(X0)), inference(ennf_transformation, [], [f41])).
fof(f41, plain, ! [X0] : (l1_lattices(X0) => (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', dt_u1_lattices)).
fof(f1052, plain, ! [X2, X3] : (l3_lattices(g3_lattices(u1_struct_0(X2), X3, u1_lattices(X2))) | ~ v1_funct_1(u1_lattices(X2)) | ~ m1_relset_1(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | ~ v1_funct_2(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | ~ v1_funct_1(X3) | ~ l1_lattices(X2)), inference(subsumption_resolution, [], [f1046, f756])).
fof(f756, plain, ! [X0] : (m1_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(resolution, [], [f387, f530])).
fof(f387, plain, ! [X0] : (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f175])).
fof(f1046, plain, ! [X2, X3] : (~ m1_relset_1(u1_lattices(X2), k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | l3_lattices(g3_lattices(u1_struct_0(X2), X3, u1_lattices(X2))) | ~ v1_funct_1(u1_lattices(X2)) | ~ m1_relset_1(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | ~ v1_funct_2(X3, k2_zfmisc_1(u1_struct_0(X2), u1_struct_0(X2)), u1_struct_0(X2)) | ~ v1_funct_1(X3) | ~ l1_lattices(X2)), inference(resolution, [], [f364, f386])).
fof(f386, plain, ! [X0] : (v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f175])).
fof(f364, plain, ! [X2, X0, X1] : (~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | l3_lattices(g3_lattices(X0, X1, X2)) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f153])).
fof(f153, plain, ! [X0, X1, X2] : ((l3_lattices(g3_lattices(X0, X1, X2)) & v3_lattices(g3_lattices(X0, X1, X2))) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(flattening, [], [f152])).
fof(f152, plain, ! [X0, X1, X2] : ((l3_lattices(g3_lattices(X0, X1, X2)) & v3_lattices(g3_lattices(X0, X1, X2))) | (~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X2) & m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X1)) => (l3_lattices(g3_lattices(X0, X1, X2)) & v3_lattices(g3_lattices(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', dt_g3_lattices)).
fof(f2222, plain, spl37_60, inference(avatar_contradiction_clause, [], [f2221])).
fof(f2221, plain, ($false | spl37_60), inference(subsumption_resolution, [], [f2220, f330])).
fof(f2220, plain, (~ l2_conlat_1(sK3) | spl37_60), inference(subsumption_resolution, [], [f2219, f329])).
fof(f2219, plain, (v3_conlat_1(sK3) | ~ l2_conlat_1(sK3) | spl37_60), inference(resolution, [], [f1347, f882])).
fof(f882, plain, ! [X0] : (m1_relset_1(k9_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) | v3_conlat_1(X0) | ~ l2_conlat_1(X0)), inference(resolution, [], [f377, f530])).
fof(f377, plain, ! [X0] : (m2_relset_1(k9_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ! [X0] : ((m2_relset_1(k9_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_2(k9_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_1(k9_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f166])).
fof(f166, plain, ! [X0] : ((m2_relset_1(k9_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_2(k9_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_1(k9_conlat_1(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (m2_relset_1(k9_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_2(k9_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_1(k9_conlat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', dt_k9_conlat_1)).
fof(f1347, plain, (~ m1_relset_1(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | spl37_60), inference(avatar_component_clause, [], [f1345])).
fof(f1345, plain, (spl37_60 <=> m1_relset_1(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_60])])).
fof(f2198, plain, (~ spl37_59 | ~ spl37_60 | spl37_63 | ~ spl37_55 | ~ spl37_56 | ~ spl37_57 | ~ spl37_58), inference(avatar_split_clause, [], [f2197, f1337, f1333, f1329, f1325, f1357, f1345, f1341])).
fof(f1341, plain, (spl37_59 <=> v1_funct_2(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_59])])).
fof(f1325, plain, (spl37_55 <=> v1_funct_1(k10_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_55])])).
fof(f1329, plain, (spl37_56 <=> v1_funct_2(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_56])])).
fof(f1333, plain, (spl37_57 <=> m1_relset_1(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_57])])).
fof(f1337, plain, (spl37_58 <=> v1_funct_1(k9_conlat_1(sK3))), introduced(avatar_definition, [new_symbols(naming, [spl37_58])])).
fof(f2197, plain, (! [X17, X15, X16] : (~ (k11_conlat_1(sK3) = g3_lattices(X15, X16, X17)) | (k8_conlat_1(sK3) = X15) | ~ m1_relset_1(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_2(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3))) | (~ spl37_55 | ~ spl37_56 | ~ spl37_57 | ~ spl37_58)), inference(subsumption_resolution, [], [f2196, f1326])).
fof(f1326, plain, (v1_funct_1(k10_conlat_1(sK3)) | ~ spl37_55), inference(avatar_component_clause, [], [f1325])).
fof(f2196, plain, (! [X17, X15, X16] : (~ (k11_conlat_1(sK3) = g3_lattices(X15, X16, X17)) | (k8_conlat_1(sK3) = X15) | ~ m1_relset_1(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_2(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_1(k10_conlat_1(sK3))) | (~ spl37_56 | ~ spl37_57 | ~ spl37_58)), inference(subsumption_resolution, [], [f2195, f1330])).
fof(f1330, plain, (v1_funct_2(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ spl37_56), inference(avatar_component_clause, [], [f1329])).
fof(f2195, plain, (! [X17, X15, X16] : (~ (k11_conlat_1(sK3) = g3_lattices(X15, X16, X17)) | (k8_conlat_1(sK3) = X15) | ~ m1_relset_1(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_2(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_2(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_1(k10_conlat_1(sK3))) | (~ spl37_57 | ~ spl37_58)), inference(subsumption_resolution, [], [f2194, f1334])).
fof(f1334, plain, (m1_relset_1(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ spl37_57), inference(avatar_component_clause, [], [f1333])).
fof(f2194, plain, (! [X17, X15, X16] : (~ (k11_conlat_1(sK3) = g3_lattices(X15, X16, X17)) | (k8_conlat_1(sK3) = X15) | ~ m1_relset_1(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_2(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ m1_relset_1(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_2(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_1(k10_conlat_1(sK3))) | ~ spl37_58), inference(subsumption_resolution, [], [f2145, f1338])).
fof(f1338, plain, (v1_funct_1(k9_conlat_1(sK3)) | ~ spl37_58), inference(avatar_component_clause, [], [f1337])).
fof(f2145, plain, ! [X17, X15, X16] : (~ (k11_conlat_1(sK3) = g3_lattices(X15, X16, X17)) | (k8_conlat_1(sK3) = X15) | ~ m1_relset_1(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_2(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_1(k9_conlat_1(sK3)) | ~ m1_relset_1(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_2(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | ~ v1_funct_1(k10_conlat_1(sK3))), inference(superposition, [], [f465, f848])).
fof(f848, plain, (k11_conlat_1(sK3) = g3_lattices(k8_conlat_1(sK3), k10_conlat_1(sK3), k9_conlat_1(sK3))), inference(subsumption_resolution, [], [f846, f329])).
fof(f846, plain, ((k11_conlat_1(sK3) = g3_lattices(k8_conlat_1(sK3), k10_conlat_1(sK3), k9_conlat_1(sK3))) | v3_conlat_1(sK3)), inference(resolution, [], [f360, f330])).
fof(f360, plain, ! [X0] : (~ l2_conlat_1(X0) | (k11_conlat_1(X0) = g3_lattices(k8_conlat_1(X0), k10_conlat_1(X0), k9_conlat_1(X0))) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ! [X0] : ((k11_conlat_1(X0) = g3_lattices(k8_conlat_1(X0), k10_conlat_1(X0), k9_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f148])).
fof(f148, plain, ! [X0] : ((k11_conlat_1(X0) = g3_lattices(k8_conlat_1(X0), k10_conlat_1(X0), k9_conlat_1(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (k11_conlat_1(X0) = g3_lattices(k8_conlat_1(X0), k10_conlat_1(X0), k9_conlat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', d23_conlat_1)).
fof(f465, plain, ! [X4, X2, X0, X5, X3, X1] : (~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) | (X0 = X3) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f215])).
fof(f215, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(flattening, [], [f214])).
fof(f214, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | (~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1))), inference(ennf_transformation, [], [f75])).
fof(f75, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X2) & m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X1)) => ! [X3, X4, X5] : ((g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) => ((X2 = X5) & (X1 = X4) & (X0 = X3)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', free_g3_lattices)).
fof(f2193, plain, spl37_59, inference(avatar_contradiction_clause, [], [f2192])).
fof(f2192, plain, ($false | spl37_59), inference(subsumption_resolution, [], [f2191, f329])).
fof(f2191, plain, (v3_conlat_1(sK3) | spl37_59), inference(subsumption_resolution, [], [f2190, f330])).
fof(f2190, plain, (~ l2_conlat_1(sK3) | v3_conlat_1(sK3) | spl37_59), inference(resolution, [], [f1343, f376])).
fof(f376, plain, ! [X0] : (v1_funct_2(k9_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f167])).
fof(f1343, plain, (~ v1_funct_2(k9_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | spl37_59), inference(avatar_component_clause, [], [f1341])).
fof(f2118, plain, ~ spl37_68, inference(avatar_contradiction_clause, [], [f2117])).
fof(f2117, plain, ($false | ~ spl37_68), inference(subsumption_resolution, [], [f2116, f329])).
fof(f2116, plain, (v3_conlat_1(sK3) | ~ spl37_68), inference(subsumption_resolution, [], [f2108, f330])).
fof(f2108, plain, (~ l2_conlat_1(sK3) | v3_conlat_1(sK3) | ~ spl37_68), inference(resolution, [], [f2018, f681])).
fof(f681, plain, ! [X0] : (~ v1_xboole_0(k8_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(duplicate_literal_removal, [], [f678])).
fof(f678, plain, ! [X0] : (~ v1_xboole_0(k8_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(resolution, [], [f383, f374])).
fof(f374, plain, ! [X0] : (m1_conlat_1(k8_conlat_1(X0), X0) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ! [X0] : (m1_conlat_1(k8_conlat_1(X0), X0) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f164])).
fof(f164, plain, ! [X0] : (m1_conlat_1(k8_conlat_1(X0), X0) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => m1_conlat_1(k8_conlat_1(X0), X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', dt_k8_conlat_1)).
fof(f383, plain, ! [X0, X1] : (~ m1_conlat_1(X1, X0) | ~ v1_xboole_0(X1) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f173])).
fof(f173, plain, ! [X0] : (! [X1] : (~ v1_xboole_0(X1) | ~ m1_conlat_1(X1, X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f172])).
fof(f172, plain, ! [X0] : (! [X1] : (~ v1_xboole_0(X1) | ~ m1_conlat_1(X1, X0)) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => ! [X1] : (m1_conlat_1(X1, X0) => ~ v1_xboole_0(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', dt_m1_conlat_1)).
fof(f2018, plain, (v1_xboole_0(k8_conlat_1(sK3)) | ~ spl37_68), inference(avatar_component_clause, [], [f2016])).
fof(f2078, plain, spl37_58, inference(avatar_contradiction_clause, [], [f2077])).
fof(f2077, plain, ($false | spl37_58), inference(subsumption_resolution, [], [f2076, f329])).
fof(f2076, plain, (v3_conlat_1(sK3) | spl37_58), inference(subsumption_resolution, [], [f2075, f330])).
fof(f2075, plain, (~ l2_conlat_1(sK3) | v3_conlat_1(sK3) | spl37_58), inference(resolution, [], [f1339, f375])).
fof(f375, plain, ! [X0] : (v1_funct_1(k9_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f167])).
fof(f1339, plain, (~ v1_funct_1(k9_conlat_1(sK3)) | spl37_58), inference(avatar_component_clause, [], [f1337])).
fof(f2013, plain, spl37_57, inference(avatar_contradiction_clause, [], [f2012])).
fof(f2012, plain, ($false | spl37_57), inference(subsumption_resolution, [], [f2011, f330])).
fof(f2011, plain, (~ l2_conlat_1(sK3) | spl37_57), inference(subsumption_resolution, [], [f2010, f329])).
fof(f2010, plain, (v3_conlat_1(sK3) | ~ l2_conlat_1(sK3) | spl37_57), inference(resolution, [], [f1335, f866])).
fof(f866, plain, ! [X0] : (m1_relset_1(k10_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) | v3_conlat_1(X0) | ~ l2_conlat_1(X0)), inference(resolution, [], [f367, f530])).
fof(f367, plain, ! [X0] : (m2_relset_1(k10_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ! [X0] : ((m2_relset_1(k10_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_2(k10_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_1(k10_conlat_1(X0))) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(flattening, [], [f154])).
fof(f154, plain, ! [X0] : ((m2_relset_1(k10_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_2(k10_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_1(k10_conlat_1(X0))) | (~ l2_conlat_1(X0) | v3_conlat_1(X0))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ! [X0] : ((l2_conlat_1(X0) & ~ v3_conlat_1(X0)) => (m2_relset_1(k10_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_2(k10_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) & v1_funct_1(k10_conlat_1(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT338+1.p', dt_k10_conlat_1)).
fof(f1335, plain, (~ m1_relset_1(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | spl37_57), inference(avatar_component_clause, [], [f1333])).
fof(f1972, plain, spl37_56, inference(avatar_contradiction_clause, [], [f1971])).
fof(f1971, plain, ($false | spl37_56), inference(subsumption_resolution, [], [f1970, f329])).
fof(f1970, plain, (v3_conlat_1(sK3) | spl37_56), inference(subsumption_resolution, [], [f1969, f330])).
fof(f1969, plain, (~ l2_conlat_1(sK3) | v3_conlat_1(sK3) | spl37_56), inference(resolution, [], [f1331, f366])).
fof(f366, plain, ! [X0] : (v1_funct_2(k10_conlat_1(X0), k2_zfmisc_1(k8_conlat_1(X0), k8_conlat_1(X0)), k8_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f155])).
fof(f1331, plain, (~ v1_funct_2(k10_conlat_1(sK3), k2_zfmisc_1(k8_conlat_1(sK3), k8_conlat_1(sK3)), k8_conlat_1(sK3)) | spl37_56), inference(avatar_component_clause, [], [f1329])).
fof(f1363, plain, spl37_55, inference(avatar_contradiction_clause, [], [f1362])).
fof(f1362, plain, ($false | spl37_55), inference(subsumption_resolution, [], [f1361, f329])).
fof(f1361, plain, (v3_conlat_1(sK3) | spl37_55), inference(subsumption_resolution, [], [f1360, f330])).
fof(f1360, plain, (~ l2_conlat_1(sK3) | v3_conlat_1(sK3) | spl37_55), inference(resolution, [], [f1327, f365])).
fof(f365, plain, ! [X0] : (v1_funct_1(k10_conlat_1(X0)) | ~ l2_conlat_1(X0) | v3_conlat_1(X0)), inference(cnf_transformation, [], [f155])).
fof(f1327, plain, (~ v1_funct_1(k10_conlat_1(sK3)) | spl37_55), inference(avatar_component_clause, [], [f1325])).
fof(f573, plain, (spl37_1 | ~ spl37_2 | ~ spl37_3 | spl37_4 | ~ spl37_5 | ~ spl37_6), inference(avatar_split_clause, [], [f332, f570, f566, f562, f558, f554, f550])).
fof(f332, plain, (~ l3_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ v9_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | v7_conlat_1(k15_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ l3_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3) | ~ v9_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3) | v7_conlat_1(k16_lattice3(k11_conlat_1(sK3), sK4), sK3)), inference(cnf_transformation, [], [f255])).