fof(f13471, plain, $false, inference(avatar_sat_refutation, [], [f1342, f2310, f2870, f2921, f3321, f3537, f6616, f6634, f6652, f6809, f13407, f13429, f13456, f13470])).
fof(f13470, plain, ~ spl42_88, inference(avatar_contradiction_clause, [], [f13469])).
fof(f13469, plain, ($false | ~ spl42_88), inference(subsumption_resolution, [], [f13468, f514])).
fof(f514, plain, v2_orders_2(sK9), inference(cnf_transformation, [], [f438])).
fof(f438, plain, ((~ v4_waybel_0(k2_yellow_2(sK9, sK9, sK10), sK9) & v1_waybel34(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) & m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v7_waybel_1(sK10, sK9) & v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(sK10)) & l1_orders_2(sK9) & v3_waybel_3(sK9) & v3_lattice3(sK9) & v2_lattice3(sK9) & v1_lattice3(sK9) & v4_orders_2(sK9) & v3_orders_2(sK9) & v2_orders_2(sK9)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10])], [f222, f437, f436])).
fof(f436, plain, (? [X0] : (? [X1] : (~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) & v1_waybel34(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_waybel_3(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (? [X1] : (~ v4_waybel_0(k2_yellow_2(sK9, sK9, X1), sK9) & v1_waybel34(k2_waybel_1(sK9, sK9, X1), sK9, k2_yellow_2(sK9, sK9, X1)) & m2_relset_1(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v7_waybel_1(X1, sK9) & v1_funct_2(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(X1)) & l1_orders_2(sK9) & v3_waybel_3(sK9) & v3_lattice3(sK9) & v2_lattice3(sK9) & v1_lattice3(sK9) & v4_orders_2(sK9) & v3_orders_2(sK9) & v2_orders_2(sK9))), introduced(choice_axiom, [])).
fof(f437, plain, (? [X1] : (~ v4_waybel_0(k2_yellow_2(sK9, sK9, X1), sK9) & v1_waybel34(k2_waybel_1(sK9, sK9, X1), sK9, k2_yellow_2(sK9, sK9, X1)) & m2_relset_1(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v7_waybel_1(X1, sK9) & v1_funct_2(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(X1)) => (~ v4_waybel_0(k2_yellow_2(sK9, sK9, sK10), sK9) & v1_waybel34(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) & m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v7_waybel_1(sK10, sK9) & v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(sK10))), introduced(choice_axiom, [])).
fof(f222, plain, ? [X0] : (? [X1] : (~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) & v1_waybel34(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_waybel_3(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(flattening, [], [f221])).
fof(f221, plain, ? [X0] : (? [X1] : ((~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) & v1_waybel34(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) & (m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1))) & (l1_orders_2(X0) & v3_waybel_3(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_waybel_3(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v1_waybel34(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) => v4_waybel_0(k2_yellow_2(X0, X0, X1), X0)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_waybel_3(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v1_waybel34(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) => v4_waybel_0(k2_yellow_2(X0, X0, X1), X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', t43_waybel34)).
fof(f13468, plain, (~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13467, f515])).
fof(f515, plain, v3_orders_2(sK9), inference(cnf_transformation, [], [f438])).
fof(f13467, plain, (~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13466, f516])).
fof(f516, plain, v4_orders_2(sK9), inference(cnf_transformation, [], [f438])).
fof(f13466, plain, (~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13465, f517])).
fof(f517, plain, v1_lattice3(sK9), inference(cnf_transformation, [], [f438])).
fof(f13465, plain, (~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13464, f518])).
fof(f518, plain, v2_lattice3(sK9), inference(cnf_transformation, [], [f438])).
fof(f13464, plain, (~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13463, f519])).
fof(f519, plain, v3_lattice3(sK9), inference(cnf_transformation, [], [f438])).
fof(f13463, plain, (~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13462, f521])).
fof(f521, plain, l1_orders_2(sK9), inference(cnf_transformation, [], [f438])).
fof(f13462, plain, (~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13461, f522])).
fof(f522, plain, v1_funct_1(sK10), inference(cnf_transformation, [], [f438])).
fof(f13461, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13460, f523])).
fof(f523, plain, v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)), inference(cnf_transformation, [], [f438])).
fof(f13460, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13459, f524])).
fof(f524, plain, v7_waybel_1(sK10, sK9), inference(cnf_transformation, [], [f438])).
fof(f13459, plain, (~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13458, f525])).
fof(f525, plain, m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)), inference(cnf_transformation, [], [f438])).
fof(f13458, plain, (~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(subsumption_resolution, [], [f13457, f527])).
fof(f527, plain, ~ v4_waybel_0(k2_yellow_2(sK9, sK9, sK10), sK9), inference(cnf_transformation, [], [f438])).
fof(f13457, plain, (v4_waybel_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ spl42_88), inference(resolution, [], [f6808, f1046])).
fof(f1046, plain, ! [X0, X1] : (~ v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) | v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f513])).
fof(f513, plain, ! [X0] : (! [X1] : (((v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) | ~ v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)) & (v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) | ~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0))) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(nnf_transformation, [], [f411])).
fof(f411, plain, ! [X0] : (! [X1] : ((v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f410])).
fof(f410, plain, ! [X0] : (! [X1] : ((v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)) | (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f144])).
fof(f144, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', t40_waybel34)).
fof(f6808, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ spl42_88), inference(avatar_component_clause, [], [f6806])).
fof(f6806, plain, (spl42_88 <=> v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9)), introduced(avatar_definition, [new_symbols(naming, [spl42_88])])).
fof(f13456, plain, spl42_75, inference(avatar_contradiction_clause, [], [f13455])).
fof(f13455, plain, ($false | spl42_75), inference(subsumption_resolution, [], [f13454, f514])).
fof(f13454, plain, (~ v2_orders_2(sK9) | spl42_75), inference(subsumption_resolution, [], [f13453, f515])).
fof(f13453, plain, (~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_75), inference(subsumption_resolution, [], [f13452, f516])).
fof(f13452, plain, (~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_75), inference(subsumption_resolution, [], [f13451, f517])).
fof(f13451, plain, (~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_75), inference(subsumption_resolution, [], [f13450, f518])).
fof(f13450, plain, (~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_75), inference(subsumption_resolution, [], [f13449, f519])).
fof(f13449, plain, (~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_75), inference(subsumption_resolution, [], [f13448, f521])).
fof(f13448, plain, (~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_75), inference(subsumption_resolution, [], [f13447, f522])).
fof(f13447, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_75), inference(subsumption_resolution, [], [f13446, f523])).
fof(f13446, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_75), inference(subsumption_resolution, [], [f13445, f524])).
fof(f13445, plain, (~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_75), inference(subsumption_resolution, [], [f13444, f525])).
fof(f13444, plain, (~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_75), inference(resolution, [], [f6547, f1040])).
fof(f1040, plain, ! [X0, X1] : (v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f408])).
fof(f408, plain, ! [X0] : (! [X1] : (((k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) & v17_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f407])).
fof(f407, plain, ! [X0] : (! [X1] : (((k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) & v17_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) | (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f142])).
fof(f142, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => ((k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) & v17_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', t39_waybel34)).
fof(f6547, plain, (~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | spl42_75), inference(avatar_component_clause, [], [f6545])).
fof(f6545, plain, (spl42_75 <=> v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl42_75])])).
fof(f13429, plain, (spl42_5 | spl42_76), inference(avatar_contradiction_clause, [], [f13428])).
fof(f13428, plain, ($false | (spl42_5 | spl42_76)), inference(subsumption_resolution, [], [f13427, f1395])).
fof(f1395, plain, m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)), inference(resolution, [], [f1034, f525])).
fof(f1034, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f512])).
fof(f512, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f137])).
fof(f137, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', redefinition_m2_relset_1)).
fof(f13427, plain, (~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | (spl42_5 | spl42_76)), inference(subsumption_resolution, [], [f13426, f1247])).
fof(f1247, plain, (~ v3_struct_0(sK9) | spl42_5), inference(avatar_component_clause, [], [f1246])).
fof(f1246, plain, (spl42_5 <=> v3_struct_0(sK9)), introduced(avatar_definition, [new_symbols(naming, [spl42_5])])).
fof(f13426, plain, (v3_struct_0(sK9) | ~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | spl42_76), inference(subsumption_resolution, [], [f13425, f521])).
fof(f13425, plain, (~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | spl42_76), inference(subsumption_resolution, [], [f13424, f522])).
fof(f13424, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | spl42_76), inference(subsumption_resolution, [], [f13423, f523])).
fof(f13423, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | spl42_76), inference(duplicate_literal_removal, [], [f13422])).
fof(f13422, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | spl42_76), inference(resolution, [], [f6551, f3388])).
fof(f3388, plain, ! [X10, X11, X9] : (m1_relset_1(k2_waybel_1(X10, X11, X9), u1_struct_0(X10), u1_struct_0(k2_yellow_2(X10, X11, X9))) | ~ v1_funct_2(X9, u1_struct_0(X10), u1_struct_0(X11)) | ~ v1_funct_1(X9) | ~ l1_orders_2(X11) | v3_struct_0(X11) | ~ l1_orders_2(X10) | v3_struct_0(X10) | ~ m1_relset_1(X9, u1_struct_0(X10), u1_struct_0(X11))), inference(resolution, [], [f689, f1034])).
fof(f689, plain, ! [X2, X0, X1] : (m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f314])).
fof(f314, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f313])).
fof(f313, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f64])).
fof(f64, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', dt_k2_waybel_1)).
fof(f6551, plain, (~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | spl42_76), inference(avatar_component_clause, [], [f6549])).
fof(f6549, plain, (spl42_76 <=> m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_76])])).
fof(f13407, plain, (spl42_5 | spl42_74), inference(avatar_contradiction_clause, [], [f13406])).
fof(f13406, plain, ($false | (spl42_5 | spl42_74)), inference(subsumption_resolution, [], [f13405, f1247])).
fof(f13405, plain, (v3_struct_0(sK9) | spl42_74), inference(subsumption_resolution, [], [f13404, f521])).
fof(f13404, plain, (~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl42_74), inference(subsumption_resolution, [], [f13403, f522])).
fof(f13403, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl42_74), inference(subsumption_resolution, [], [f13402, f523])).
fof(f13402, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl42_74), inference(subsumption_resolution, [], [f13401, f1395])).
fof(f13401, plain, (~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl42_74), inference(duplicate_literal_removal, [], [f13400])).
fof(f13400, plain, (~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl42_74), inference(resolution, [], [f6543, f688])).
fof(f688, plain, ! [X2, X0, X1] : (v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f314])).
fof(f6543, plain, (~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | spl42_74), inference(avatar_component_clause, [], [f6541])).
fof(f6541, plain, (spl42_74 <=> v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))), introduced(avatar_definition, [new_symbols(naming, [spl42_74])])).
fof(f6809, plain, (~ spl42_82 | ~ spl42_80 | ~ spl42_81 | spl42_88 | spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_37 | ~ spl42_38), inference(avatar_split_clause, [], [f6804, f2924, f2918, f2905, f2453, f1246, f6806, f6631, f6613, f6649])).
fof(f6649, plain, (spl42_82 <=> v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9))), introduced(avatar_definition, [new_symbols(naming, [spl42_82])])).
fof(f6613, plain, (spl42_80 <=> v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9)), introduced(avatar_definition, [new_symbols(naming, [spl42_80])])).
fof(f6631, plain, (spl42_81 <=> m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9))), introduced(avatar_definition, [new_symbols(naming, [spl42_81])])).
fof(f2453, plain, (spl42_28 <=> l1_orders_2(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl42_28])])).
fof(f2905, plain, (spl42_35 <=> v3_lattice3(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl42_35])])).
fof(f2918, plain, (spl42_37 <=> v1_lattice3(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl42_37])])).
fof(f2924, plain, (spl42_38 <=> v2_lattice3(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl42_38])])).
fof(f6804, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | (spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6803, f2760])).
fof(f2760, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f2759, f514])).
fof(f2759, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f2758, f521])).
fof(f2758, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f2751, f2734])).
fof(f2734, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl42_5), inference(subsumption_resolution, [], [f2733, f1247])).
fof(f2733, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2732, f521])).
fof(f2732, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2731, f522])).
fof(f2731, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2730, f523])).
fof(f2730, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2713])).
fof(f2713, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f692, f1395])).
fof(f692, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f316])).
fof(f316, plain, ! [X0, X1, X2] : ((m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f315])).
fof(f315, plain, ! [X0, X1, X2] : ((m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f65])).
fof(f65, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', dt_k2_yellow_2)).
fof(f2751, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_5), inference(resolution, [], [f2698, f660])).
fof(f660, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v2_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f291])).
fof(f291, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v2_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f290])).
fof(f290, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v2_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f50])).
fof(f50, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v2_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', cc6_yellow_0)).
fof(f2698, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl42_5), inference(subsumption_resolution, [], [f2697, f1247])).
fof(f2697, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2696, f521])).
fof(f2696, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2695, f522])).
fof(f2695, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2694, f523])).
fof(f2694, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2677])).
fof(f2677, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f691, f1395])).
fof(f691, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f316])).
fof(f6803, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6802, f2757])).
fof(f2757, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f2756, f515])).
fof(f2756, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f2755, f521])).
fof(f2755, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f2750, f2734])).
fof(f2750, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v3_orders_2(sK9) | spl42_5), inference(resolution, [], [f2698, f665])).
fof(f665, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v3_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v3_orders_2(X0)), inference(cnf_transformation, [], [f295])).
fof(f295, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v3_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v3_orders_2(X0)), inference(flattening, [], [f294])).
fof(f294, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v3_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v3_orders_2(X0))), inference(ennf_transformation, [], [f52])).
fof(f52, plain, ! [X0] : ((l1_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v3_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', cc7_yellow_0)).
fof(f6802, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6801, f2754])).
fof(f2754, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f2753, f516])).
fof(f2753, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f2752, f521])).
fof(f2752, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f2749, f2734])).
fof(f2749, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | spl42_5), inference(resolution, [], [f2698, f670])).
fof(f670, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v4_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0)), inference(cnf_transformation, [], [f299])).
fof(f299, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v4_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0)), inference(flattening, [], [f298])).
fof(f298, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v4_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v4_orders_2(X0))), inference(ennf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : ((l1_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v4_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', cc8_yellow_0)).
fof(f6801, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6800, f2920])).
fof(f2920, plain, (v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl42_37), inference(avatar_component_clause, [], [f2918])).
fof(f6800, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_38)), inference(subsumption_resolution, [], [f6799, f2926])).
fof(f2926, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl42_38), inference(avatar_component_clause, [], [f2924])).
fof(f6799, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_35)), inference(subsumption_resolution, [], [f6798, f2906])).
fof(f2906, plain, (v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl42_35), inference(avatar_component_clause, [], [f2905])).
fof(f6798, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28)), inference(subsumption_resolution, [], [f6797, f2454])).
fof(f2454, plain, (l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ spl42_28), inference(avatar_component_clause, [], [f2453])).
fof(f6797, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6796, f514])).
fof(f6796, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6795, f515])).
fof(f6795, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6794, f516])).
fof(f6794, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6793, f517])).
fof(f6793, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6792, f518])).
fof(f6792, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6791, f519])).
fof(f6791, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6790, f520])).
fof(f520, plain, v3_waybel_3(sK9), inference(cnf_transformation, [], [f438])).
fof(f6790, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_waybel_3(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6789, f521])).
fof(f6789, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ l1_orders_2(sK9) | ~ v3_waybel_3(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6788, f2482])).
fof(f2482, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f2481, f1247])).
fof(f2481, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2480, f521])).
fof(f2480, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2479, f522])).
fof(f2479, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2478, f523])).
fof(f2478, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2461])).
fof(f2461, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f693, f1395])).
fof(f693, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v1_funct_1(k3_waybel_1(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f318])).
fof(f318, plain, ! [X0, X1, X2] : ((m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f317])).
fof(f317, plain, ! [X0, X1, X2] : ((m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', dt_k3_waybel_1)).
fof(f6788, plain, (v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_waybel_3(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), inference(subsumption_resolution, [], [f6674, f526])).
fof(f526, plain, v1_waybel34(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)), inference(cnf_transformation, [], [f438])).
fof(f6674, plain, (~ v1_waybel34(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_waybel_3(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), inference(superposition, [], [f1038, f4116])).
fof(f4116, plain, (k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))), inference(subsumption_resolution, [], [f4115, f514])).
fof(f4115, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4114, f515])).
fof(f4114, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4113, f516])).
fof(f4113, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4112, f517])).
fof(f4112, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4111, f518])).
fof(f4111, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4110, f519])).
fof(f4110, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4109, f521])).
fof(f4109, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4108, f522])).
fof(f4108, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4107, f523])).
fof(f4107, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4100, f524])).
fof(f4100, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(resolution, [], [f1043, f525])).
fof(f1043, plain, ! [X0, X1] : (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | (k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f408])).
fof(f1038, plain, ! [X2, X0, X1] : (~ v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | v22_waybel_0(X2, X0, X1) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_waybel_3(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f404])).
fof(f404, plain, ! [X0] : (! [X1] : (! [X2] : (v22_waybel_0(X2, X0, X1) | ~ v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_waybel_3(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f403])).
fof(f403, plain, ! [X0] : (! [X1] : (! [X2] : ((v22_waybel_0(X2, X0, X1) | ~ v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0)) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v3_waybel_3(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f140])).
fof(f140, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_waybel_3(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => (v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) => v22_waybel_0(X2, X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', t25_waybel34)).
fof(f6652, plain, (~ spl42_74 | ~ spl42_76 | spl42_82 | spl42_5 | ~ spl42_28 | ~ spl42_37 | ~ spl42_38), inference(avatar_split_clause, [], [f6647, f2924, f2918, f2453, f1246, f6649, f6549, f6541])).
fof(f6647, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | (spl42_5 | ~ spl42_28 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6646, f2760])).
fof(f6646, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6645, f2757])).
fof(f6645, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6644, f2754])).
fof(f6644, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6643, f2920])).
fof(f6643, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_38)), inference(subsumption_resolution, [], [f6642, f2926])).
fof(f6642, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28)), inference(subsumption_resolution, [], [f6641, f2454])).
fof(f6641, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6640, f514])).
fof(f6640, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6639, f515])).
fof(f6639, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6638, f516])).
fof(f6638, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6637, f517])).
fof(f6637, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6636, f518])).
fof(f6636, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6635, f521])).
fof(f6635, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6524, f2400])).
fof(f2400, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f2399, f1247])).
fof(f2399, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2398, f521])).
fof(f2398, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2397, f522])).
fof(f2397, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2396, f523])).
fof(f2396, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2379])).
fof(f2379, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f687, f1395])).
fof(f687, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v1_funct_1(k2_waybel_1(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f314])).
fof(f6524, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), inference(superposition, [], [f685, f4097])).
fof(f4097, plain, (k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))), inference(subsumption_resolution, [], [f4096, f514])).
fof(f4096, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4095, f515])).
fof(f4095, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4094, f516])).
fof(f4094, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4093, f517])).
fof(f4093, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4092, f518])).
fof(f4092, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4091, f519])).
fof(f4091, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4090, f521])).
fof(f4090, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4089, f522])).
fof(f4089, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4088, f523])).
fof(f4088, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f4081, f524])).
fof(f4081, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(resolution, [], [f1042, f525])).
fof(f1042, plain, ! [X0, X1] : (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f408])).
fof(f685, plain, ! [X2, X0, X1] : (v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f312])).
fof(f312, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f311])).
fof(f311, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f63])).
fof(f63, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', dt_k2_waybel34)).
fof(f6634, plain, (~ spl42_74 | ~ spl42_76 | spl42_81 | spl42_5 | ~ spl42_28 | ~ spl42_37 | ~ spl42_38), inference(avatar_split_clause, [], [f6629, f2924, f2918, f2453, f1246, f6631, f6549, f6541])).
fof(f6629, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | (spl42_5 | ~ spl42_28 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6628, f2760])).
fof(f6628, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6627, f2757])).
fof(f6627, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6626, f2754])).
fof(f6626, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6625, f2920])).
fof(f6625, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_38)), inference(subsumption_resolution, [], [f6624, f2926])).
fof(f6624, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28)), inference(subsumption_resolution, [], [f6623, f2454])).
fof(f6623, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6622, f514])).
fof(f6622, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6621, f515])).
fof(f6621, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6620, f516])).
fof(f6620, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6619, f517])).
fof(f6619, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6618, f518])).
fof(f6618, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6617, f521])).
fof(f6617, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6523, f2400])).
fof(f6523, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), inference(superposition, [], [f686, f4097])).
fof(f686, plain, ! [X2, X0, X1] : (m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f312])).
fof(f6616, plain, (~ spl42_74 | ~ spl42_75 | ~ spl42_76 | spl42_80 | spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_37 | ~ spl42_38), inference(avatar_split_clause, [], [f6611, f2924, f2918, f2905, f2453, f1246, f6613, f6549, f6545, f6541])).
fof(f6611, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | (spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6610, f2760])).
fof(f6610, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6609, f2757])).
fof(f6609, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6608, f2754])).
fof(f6608, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_37 | ~ spl42_38)), inference(subsumption_resolution, [], [f6607, f2920])).
fof(f6607, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_35 | ~ spl42_38)), inference(subsumption_resolution, [], [f6606, f2926])).
fof(f6606, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28 | ~ spl42_35)), inference(subsumption_resolution, [], [f6605, f2906])).
fof(f6605, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28)), inference(subsumption_resolution, [], [f6604, f2454])).
fof(f6604, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6603, f514])).
fof(f6603, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6602, f515])).
fof(f6602, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6601, f516])).
fof(f6601, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6600, f517])).
fof(f6600, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6599, f518])).
fof(f6599, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6598, f519])).
fof(f6598, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6597, f521])).
fof(f6597, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f6522, f2400])).
fof(f6522, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), inference(superposition, [], [f767, f4097])).
fof(f767, plain, ! [X2, X0, X1] : (v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f343])).
fof(f343, plain, ! [X0, X1, X2] : ((v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f342])).
fof(f342, plain, ! [X0, X1, X2] : ((v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f203])).
fof(f203, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f172])).
fof(f172, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_partfun1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f92])).
fof(f92, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v21_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_partfun1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', fc2_waybel34)).
fof(f3537, plain, (~ spl42_27 | spl42_35), inference(avatar_contradiction_clause, [], [f3536])).
fof(f3536, plain, ($false | (~ spl42_27 | spl42_35)), inference(subsumption_resolution, [], [f3535, f514])).
fof(f3535, plain, (~ v2_orders_2(sK9) | (~ spl42_27 | spl42_35)), inference(subsumption_resolution, [], [f3534, f515])).
fof(f3534, plain, (~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl42_27 | spl42_35)), inference(subsumption_resolution, [], [f3533, f516])).
fof(f3533, plain, (~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl42_27 | spl42_35)), inference(subsumption_resolution, [], [f3532, f517])).
fof(f3532, plain, (~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl42_27 | spl42_35)), inference(subsumption_resolution, [], [f3531, f518])).
fof(f3531, plain, (~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl42_27 | spl42_35)), inference(subsumption_resolution, [], [f3530, f519])).
fof(f3530, plain, (~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl42_27 | spl42_35)), inference(subsumption_resolution, [], [f3529, f521])).
fof(f3529, plain, (~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl42_27 | spl42_35)), inference(subsumption_resolution, [], [f3528, f522])).
fof(f3528, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl42_27 | spl42_35)), inference(subsumption_resolution, [], [f3527, f523])).
fof(f3527, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl42_27 | spl42_35)), inference(subsumption_resolution, [], [f3526, f2201])).
fof(f2201, plain, (v6_waybel_1(sK10, sK9) | ~ spl42_27), inference(avatar_component_clause, [], [f2200])).
fof(f2200, plain, (spl42_27 <=> v6_waybel_1(sK10, sK9)), introduced(avatar_definition, [new_symbols(naming, [spl42_27])])).
fof(f3526, plain, (~ v6_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl42_35), inference(subsumption_resolution, [], [f3508, f3006])).
fof(f3006, plain, (~ sP2(sK10, sK9) | spl42_35), inference(resolution, [], [f2907, f731])).
fof(f731, plain, ! [X0, X1] : (v3_lattice3(k2_yellow_2(X1, X1, X0)) | ~ sP2(X0, X1)), inference(cnf_transformation, [], [f454])).
fof(f454, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X1, X1, X0)) & v2_lattice3(k2_yellow_2(X1, X1, X0)) & v1_lattice3(k2_yellow_2(X1, X1, X0)) & v25_waybel_0(k2_yellow_2(X1, X1, X0)) & v24_waybel_0(k2_yellow_2(X1, X1, X0)) & v4_yellow_0(k2_yellow_2(X1, X1, X0), X1) & v3_yellow_0(k2_yellow_2(X1, X1, X0)) & v2_yellow_0(k2_yellow_2(X1, X1, X0)) & v1_yellow_0(k2_yellow_2(X1, X1, X0)) & v4_orders_2(k2_yellow_2(X1, X1, X0)) & v3_orders_2(k2_yellow_2(X1, X1, X0)) & v2_orders_2(k2_yellow_2(X1, X1, X0)) & v1_orders_2(k2_yellow_2(X1, X1, X0)) & ~ v3_struct_0(k2_yellow_2(X1, X1, X0))) | ~ sP2(X0, X1)), inference(rectify, [], [f453])).
fof(f453, plain, ! [X1, X0] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ sP2(X1, X0)), inference(nnf_transformation, [], [f422])).
fof(f422, plain, ! [X1, X0] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ sP2(X1, X0)), inference(usedef, [], [e422])).
fof(e422, plain, ! [X1, X0] : (sP2(X1, X0) <=> (v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f2907, plain, (~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | spl42_35), inference(avatar_component_clause, [], [f2905])).
fof(f3508, plain, (sP2(sK10, sK9) | ~ v6_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(resolution, [], [f732, f1395])).
fof(f732, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | sP2(X1, X0) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f423])).
fof(f423, plain, ! [X0, X1] : (sP2(X1, X0) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(definition_folding, [], [f329, e422])).
fof(f329, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f328])).
fof(f328, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f85])).
fof(f85, plain, ! [X0, X1] : ((m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', fc10_waybel34)).
fof(f3321, plain, (spl42_38 | spl42_5), inference(avatar_split_clause, [], [f3320, f1246, f2924])).
fof(f3320, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f3319, f515])).
fof(f3319, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f3318, f516])).
fof(f3318, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f3317, f518])).
fof(f3317, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f3316, f521])).
fof(f3316, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f3315, f2734])).
fof(f3315, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f3314, f2518])).
fof(f2518, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f2517, f1247])).
fof(f2517, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2516, f521])).
fof(f2516, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2515, f522])).
fof(f2515, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2514, f523])).
fof(f2514, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2497])).
fof(f2497, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f775, f1395])).
fof(f775, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v3_struct_0(k2_yellow_2(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f346])).
fof(f346, plain, ! [X0, X1, X2] : ((v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f345])).
fof(f345, plain, ! [X0, X1, X2] : ((v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f94])).
fof(f94, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', fc2_yellow_2)).
fof(f3314, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl42_5), inference(subsumption_resolution, [], [f3311, f2698])).
fof(f3311, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl42_5), inference(resolution, [], [f3290, f540])).
fof(f540, plain, ! [X0, X1] : (~ v5_yellow_0(X1, X0) | v2_lattice3(X1) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0)), inference(cnf_transformation, [], [f231])).
fof(f231, plain, ! [X0] : (! [X1] : ((v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v2_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1)) | ~ v5_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0)), inference(flattening, [], [f230])).
fof(f230, plain, ! [X0] : (! [X1] : (((v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v2_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1)) | (~ v5_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1))) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ! [X0] : ((l1_orders_2(X0) & v2_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => ((v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & ~ v3_struct_0(X1)) => (v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v2_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', cc11_yellow_0)).
fof(f3290, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl42_5), inference(subsumption_resolution, [], [f3289, f1247])).
fof(f3289, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3288, f514])).
fof(f3288, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3287, f515])).
fof(f3287, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3286, f516])).
fof(f3286, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3285, f521])).
fof(f3285, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3284, f522])).
fof(f3284, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3283, f523])).
fof(f3283, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3265, f524])).
fof(f3265, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f716, f1395])).
fof(f716, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f327])).
fof(f327, plain, ! [X0, X1] : ((v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f326])).
fof(f326, plain, ! [X0, X1] : ((v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f180])).
fof(f180, plain, ! [X0, X1] : ((m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1) & l1_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), inference(pure_predicate_removal, [], [f84])).
fof(f84, plain, ! [X0, X1] : ((m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1) & l1_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v3_waybel_0(k2_yellow_2(X0, X0, X1), X0) & v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', fc10_waybel10)).
fof(f2921, plain, (spl42_37 | ~ spl42_35 | spl42_5 | ~ spl42_28), inference(avatar_split_clause, [], [f2916, f2453, f1246, f2905, f2918])).
fof(f2916, plain, (~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | (spl42_5 | ~ spl42_28)), inference(subsumption_resolution, [], [f2888, f2518])).
fof(f2888, plain, (~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl42_28), inference(resolution, [], [f2454, f600])).
fof(f600, plain, ! [X0] : (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | v3_struct_0(X0) | v1_lattice3(X0)), inference(cnf_transformation, [], [f258])).
fof(f258, plain, ! [X0] : ((v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)) | ~ v3_lattice3(X0) | v3_struct_0(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f257])).
fof(f257, plain, ! [X0] : (((v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)) | (~ v3_lattice3(X0) | v3_struct_0(X0))) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ! [X0] : (l1_orders_2(X0) => ((v3_lattice3(X0) & ~ v3_struct_0(X0)) => (v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', cc1_yellow_0)).
fof(f2870, plain, (spl42_28 | spl42_5), inference(avatar_split_clause, [], [f2869, f1246, f2453])).
fof(f2869, plain, (l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl42_5), inference(subsumption_resolution, [], [f2865, f521])).
fof(f2865, plain, (l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | spl42_5), inference(resolution, [], [f2734, f701])).
fof(f701, plain, ! [X0, X1] : (~ m1_yellow_0(X1, X0) | l1_orders_2(X1) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f322])).
fof(f322, plain, ! [X0] : (! [X1] : (l1_orders_2(X1) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f74])).
fof(f74, plain, ! [X0] : (l1_orders_2(X0) => ! [X1] : (m1_yellow_0(X1, X0) => l1_orders_2(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', dt_m1_yellow_0)).
fof(f2310, plain, (spl42_27 | spl42_5), inference(avatar_split_clause, [], [f2309, f1246, f2200])).
fof(f2309, plain, (v6_waybel_1(sK10, sK9) | spl42_5), inference(subsumption_resolution, [], [f2308, f1247])).
fof(f2308, plain, (v6_waybel_1(sK10, sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2307, f521])).
fof(f2307, plain, (v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2306, f522])).
fof(f2306, plain, (~ v1_funct_1(sK10) | v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2305, f523])).
fof(f2305, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2289, f524])).
fof(f2289, plain, (~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f635, f1395])).
fof(f635, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | v6_waybel_1(X1, X0) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f274])).
fof(f274, plain, ! [X0] : (! [X1] : ((v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0))) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f273])).
fof(f273, plain, ! [X0] : (! [X1] : (((v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1)) | (~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0))) | (~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f186])).
fof(f186, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) => ((v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1))))), inference(pure_predicate_removal, [], [f41])).
fof(f41, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) => ((v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v6_waybel_1(X1, X0) & v1_partfun1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', cc3_waybel_1)).
fof(f1342, plain, ~ spl42_5, inference(avatar_contradiction_clause, [], [f1341])).
fof(f1341, plain, ($false | ~ spl42_5), inference(subsumption_resolution, [], [f1340, f521])).
fof(f1340, plain, (~ l1_orders_2(sK9) | ~ spl42_5), inference(subsumption_resolution, [], [f1336, f517])).
fof(f1336, plain, (~ v1_lattice3(sK9) | ~ l1_orders_2(sK9) | ~ spl42_5), inference(resolution, [], [f1248, f580])).
fof(f580, plain, ! [X0] : (~ v3_struct_0(X0) | ~ v1_lattice3(X0) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f246])).
fof(f246, plain, ! [X0] : (~ v3_struct_0(X0) | ~ v1_lattice3(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f245])).
fof(f245, plain, ! [X0] : ((~ v3_struct_0(X0) | ~ v1_lattice3(X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ! [X0] : (l1_orders_2(X0) => (v1_lattice3(X0) => ~ v3_struct_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT370+1.p', cc1_lattice3)).
fof(f1248, plain, (v3_struct_0(sK9) | ~ spl42_5), inference(avatar_component_clause, [], [f1246])).