fof(f4037, plain, $false, inference(avatar_sat_refutation, [], [f286, f295, f300, f332, f340, f1287, f1323, f3286, f3312, f3313, f3633, f3756, f3762, f3860, f3866, f4035])).
fof(f4035, plain, (~ spl23_4 | ~ spl23_225 | spl23_226), inference(avatar_contradiction_clause, [], [f4034])).
fof(f4034, plain, ($false | (~ spl23_4 | ~ spl23_225 | spl23_226)), inference(subsumption_resolution, [], [f4028, f3761])).
fof(f3761, plain, (~ sdtlseqdt0(sK19(sK21), sK21) | spl23_226), inference(avatar_component_clause, [], [f3759])).
fof(f3759, plain, (spl23_226 <=> sdtlseqdt0(sK19(sK21), sK21)), introduced(avatar_definition, [new_symbols(naming, [spl23_226])])).
fof(f4028, plain, (sdtlseqdt0(sK19(sK21), sK21) | (~ spl23_4 | ~ spl23_225)), inference(resolution, [], [f3755, f294])).
fof(f294, plain, (! [X1] : (~ aElementOf0(X1, xT) | sdtlseqdt0(X1, sK21)) | ~ spl23_4), inference(avatar_component_clause, [], [f293])).
fof(f293, plain, (spl23_4 <=> ! [X1] : (sdtlseqdt0(X1, sK21) | ~ aElementOf0(X1, xT))), introduced(avatar_definition, [new_symbols(naming, [spl23_4])])).
fof(f3755, plain, (aElementOf0(sK19(sK21), xT) | ~ spl23_225), inference(avatar_component_clause, [], [f3753])).
fof(f3753, plain, (spl23_225 <=> aElementOf0(sK19(sK21), xT)), introduced(avatar_definition, [new_symbols(naming, [spl23_225])])).
fof(f3866, plain, (spl23_2 | ~ spl23_223), inference(avatar_contradiction_clause, [], [f3865])).
fof(f3865, plain, ($false | (spl23_2 | ~ spl23_223)), inference(subsumption_resolution, [], [f3861, f285])).
fof(f285, plain, (~ sdtlseqdt0(xp, sK21) | spl23_2), inference(avatar_component_clause, [], [f283])).
fof(f283, plain, (spl23_2 <=> sdtlseqdt0(xp, sK21)), introduced(avatar_definition, [new_symbols(naming, [spl23_2])])).
fof(f3861, plain, (sdtlseqdt0(xp, sK21) | ~ spl23_223), inference(resolution, [], [f3744, f245])).
fof(f245, plain, ! [X2] : (~ aElementOf0(X2, xP) | sdtlseqdt0(xp, X2)), inference(cnf_transformation, [], [f141])).
fof(f141, plain, (aInfimumOfIn0(xp, xP, xU) & ! [X0] : (sdtlseqdt0(X0, xp) | (~ aLowerBoundOfIn0(X0, xP, xU) & ((~ sdtlseqdt0(X0, sK20(X0)) & aElementOf0(sK20(X0), xP)) | ~ aElementOf0(X0, xU)))) & aLowerBoundOfIn0(xp, xP, xU) & ! [X2] : (sdtlseqdt0(xp, X2) | ~ aElementOf0(X2, xP)) & aElementOf0(xp, xU) & aElementOf0(xp, xU)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK20])], [f73, f140])).
fof(f140, plain, ! [X0] : (? [X1] : (~ sdtlseqdt0(X0, X1) & aElementOf0(X1, xP)) => (~ sdtlseqdt0(X0, sK20(X0)) & aElementOf0(sK20(X0), xP))), introduced(choice_axiom, [])).
fof(f73, plain, (aInfimumOfIn0(xp, xP, xU) & ! [X0] : (sdtlseqdt0(X0, xp) | (~ aLowerBoundOfIn0(X0, xP, xU) & (? [X1] : (~ sdtlseqdt0(X0, X1) & aElementOf0(X1, xP)) | ~ aElementOf0(X0, xU)))) & aLowerBoundOfIn0(xp, xP, xU) & ! [X2] : (sdtlseqdt0(xp, X2) | ~ aElementOf0(X2, xP)) & aElementOf0(xp, xU) & aElementOf0(xp, xU)), inference(ennf_transformation, [], [f38])).
fof(f38, plain, (aInfimumOfIn0(xp, xP, xU) & ! [X0] : ((aLowerBoundOfIn0(X0, xP, xU) | (! [X1] : (aElementOf0(X1, xP) => sdtlseqdt0(X0, X1)) & aElementOf0(X0, xU))) => sdtlseqdt0(X0, xp)) & aLowerBoundOfIn0(xp, xP, xU) & ! [X2] : (aElementOf0(X2, xP) => sdtlseqdt0(xp, X2)) & aElementOf0(xp, xU) & aElementOf0(xp, xU)), inference(rectify, [], [f28])).
fof(f28, plain, (aInfimumOfIn0(xp, xP, xU) & ! [X0] : ((aLowerBoundOfIn0(X0, xP, xU) | (! [X1] : (aElementOf0(X1, xP) => sdtlseqdt0(X0, X1)) & aElementOf0(X0, xU))) => sdtlseqdt0(X0, xp)) & aLowerBoundOfIn0(xp, xP, xU) & ! [X0] : (aElementOf0(X0, xP) => sdtlseqdt0(xp, X0)) & aElementOf0(xp, xU) & aElementOf0(xp, xU)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT387+4.p', m__1261)).
fof(f3744, plain, (aElementOf0(sK21, xP) | ~ spl23_223), inference(avatar_component_clause, [], [f3743])).
fof(f3743, plain, (spl23_223 <=> aElementOf0(sK21, xP)), introduced(avatar_definition, [new_symbols(naming, [spl23_223])])).
fof(f3860, plain, (~ spl23_5 | spl23_224), inference(avatar_contradiction_clause, [], [f3859])).
fof(f3859, plain, ($false | (~ spl23_5 | spl23_224)), inference(subsumption_resolution, [], [f3858, f3589])).
fof(f3589, plain, (aElement0(sK21) | ~ spl23_5), inference(subsumption_resolution, [], [f3588, f224])).
fof(f224, plain, aSet0(xS), inference(cnf_transformation, [], [f69])).
fof(f69, plain, ((xS = cS1142(xf)) & ! [X0] : ((aElementOf0(X0, xS) | (~ aFixedPointOf0(X0, xf) & (~ (sdtlpdtrp0(xf, X0) = X0) | ~ aElementOf0(X0, szDzozmdt0(xf))))) & ((aFixedPointOf0(X0, xf) & (sdtlpdtrp0(xf, X0) = X0) & aElementOf0(X0, szDzozmdt0(xf))) | ~ aElementOf0(X0, xS))) & aSet0(xS)), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ((xS = cS1142(xf)) & ! [X0] : (((aFixedPointOf0(X0, xf) | ((sdtlpdtrp0(xf, X0) = X0) & aElementOf0(X0, szDzozmdt0(xf)))) => aElementOf0(X0, xS)) & (aElementOf0(X0, xS) => (aFixedPointOf0(X0, xf) & (sdtlpdtrp0(xf, X0) = X0) & aElementOf0(X0, szDzozmdt0(xf))))) & aSet0(xS)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT387+4.p', m__1144)).
fof(f3588, plain, (aElement0(sK21) | ~ aSet0(xS) | ~ spl23_5), inference(resolution, [], [f299, f148])).
fof(f148, plain, ! [X0, X1] : (~ aElementOf0(X1, X0) | aElement0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f42])).
fof(f42, plain, ! [X0] : (! [X1] : (aElement0(X1) | ~ aElementOf0(X1, X0)) | ~ aSet0(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aElementOf0(X1, X0) => aElement0(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT387+4.p', mEOfElem)).
fof(f299, plain, (aElementOf0(sK21, xS) | ~ spl23_5), inference(avatar_component_clause, [], [f297])).
fof(f297, plain, (spl23_5 <=> aElementOf0(sK21, xS)), introduced(avatar_definition, [new_symbols(naming, [spl23_5])])).
fof(f3858, plain, (~ aElement0(sK21) | spl23_224), inference(resolution, [], [f3748, f155])).
fof(f155, plain, ! [X0] : (sdtlseqdt0(X0, X0) | ~ aElement0(X0)), inference(cnf_transformation, [], [f45])).
fof(f45, plain, ! [X0] : (sdtlseqdt0(X0, X0) | ~ aElement0(X0)), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ! [X0] : (aElement0(X0) => sdtlseqdt0(X0, X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT387+4.p', mARefl)).
fof(f3748, plain, (~ sdtlseqdt0(sK21, sK21) | spl23_224), inference(avatar_component_clause, [], [f3747])).
fof(f3747, plain, (spl23_224 <=> sdtlseqdt0(sK21, sK21)), introduced(avatar_definition, [new_symbols(naming, [spl23_224])])).
fof(f3762, plain, (spl23_223 | ~ spl23_226 | ~ spl23_224 | ~ spl23_5), inference(avatar_split_clause, [], [f3757, f297, f3747, f3759, f3743])).
fof(f3757, plain, (~ sdtlseqdt0(sK21, sK21) | ~ sdtlseqdt0(sK19(sK21), sK21) | aElementOf0(sK21, xP) | ~ spl23_5), inference(subsumption_resolution, [], [f3726, f3585])).
fof(f3585, plain, (aElementOf0(sK21, xU) | ~ spl23_5), inference(resolution, [], [f299, f277])).
fof(f277, plain, ! [X0] : (~ aElementOf0(X0, xS) | aElementOf0(X0, xU)), inference(forward_demodulation, [], [f225, f273])).
fof(f273, plain, (xU = szDzozmdt0(xf)), inference(forward_demodulation, [], [f221, f222])).
fof(f222, plain, (xU = szRzazndt0(xf)), inference(cnf_transformation, [], [f137])).
fof(f137, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (sP2(X2) | (~ aSubsetOf0(X2, xU) & ((~ aElementOf0(sK18(X2), xU) & aElementOf0(sK18(X2), X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK18])], [f80, f136])).
fof(f136, plain, ! [X2] : (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) => (~ aElementOf0(sK18(X2), xU) & aElementOf0(sK18(X2), X2))), introduced(choice_axiom, [])).
fof(f80, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (sP2(X2) | (~ aSubsetOf0(X2, xU) & (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(definition_folding, [], [f68, e79, e78, e77])).
fof(f77, plain, ! [X2] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) | ~ sP0(X2)), inference(usedef, [], [e77])).
fof(e77, plain, ! [X2] : (sP0(X2) <=> ? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f78, plain, ! [X4, X2] : (! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU)))) | ~ sP1(X4, X2)), inference(usedef, [], [e78])).
fof(e78, plain, ! [X4, X2] : (sP1(X4, X2) <=> ! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f79, plain, ! [X2] : (? [X4] : (sP0(X2) & aInfimumOfIn0(X4, X2, xU) & sP1(X4, X2) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU)) | ~ sP2(X2)), inference(usedef, [], [e79])).
fof(e79, plain, ! [X2] : (sP2(X2) <=> ? [X4] : (sP0(X2) & aInfimumOfIn0(X4, X2, xU) & sP1(X4, X2) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f68, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (? [X4] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) & aInfimumOfIn0(X4, X2, xU) & ! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU)))) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU)) | (~ aSubsetOf0(X2, xU) & (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(flattening, [], [f67])).
fof(f67, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : ((sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1)) | (~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : (? [X4] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : (sdtlseqdt0(X5, X6) | (~ aUpperBoundOfIn0(X6, X2, xU) & (? [X7] : (~ sdtlseqdt0(X7, X6) & aElementOf0(X7, X2)) | ~ aElementOf0(X6, xU)))) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (sdtlseqdt0(X8, X5) | ~ aElementOf0(X8, X2)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) & aInfimumOfIn0(X4, X2, xU) & ! [X9] : (sdtlseqdt0(X9, X4) | (~ aLowerBoundOfIn0(X9, X2, xU) & (? [X10] : (~ sdtlseqdt0(X9, X10) & aElementOf0(X10, X2)) | ~ aElementOf0(X9, xU)))) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (sdtlseqdt0(X4, X11) | ~ aElementOf0(X11, X2)) & aElementOf0(X4, xU) & aElementOf0(X4, xU)) | (~ aSubsetOf0(X2, xU) & (? [X3] : (~ aElementOf0(X3, xU) & aElementOf0(X3, X2)) | ~ aSet0(X2)))) & aSet0(xU)), inference(ennf_transformation, [], [f36])).
fof(f36, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : ((aElementOf0(X1, szDzozmdt0(xf)) & aElementOf0(X0, szDzozmdt0(xf))) => (sdtlseqdt0(X0, X1) => sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X2] : ((aSubsetOf0(X2, xU) | (! [X3] : (aElementOf0(X3, X2) => aElementOf0(X3, xU)) & aSet0(X2))) => ? [X4] : (? [X5] : (aSupremumOfIn0(X5, X2, xU) & ! [X6] : ((aUpperBoundOfIn0(X6, X2, xU) | (! [X7] : (aElementOf0(X7, X2) => sdtlseqdt0(X7, X6)) & aElementOf0(X6, xU))) => sdtlseqdt0(X5, X6)) & aUpperBoundOfIn0(X5, X2, xU) & ! [X8] : (aElementOf0(X8, X2) => sdtlseqdt0(X8, X5)) & aElementOf0(X5, xU) & aElementOf0(X5, xU)) & aInfimumOfIn0(X4, X2, xU) & ! [X9] : ((aLowerBoundOfIn0(X9, X2, xU) | (! [X10] : (aElementOf0(X10, X2) => sdtlseqdt0(X9, X10)) & aElementOf0(X9, xU))) => sdtlseqdt0(X9, X4)) & aLowerBoundOfIn0(X4, X2, xU) & ! [X11] : (aElementOf0(X11, X2) => sdtlseqdt0(X4, X11)) & aElementOf0(X4, xU) & aElementOf0(X4, xU))) & aSet0(xU)), inference(rectify, [], [f24])).
fof(f24, plain, (isOn0(xf, xU) & (xU = szRzazndt0(xf)) & (szDzozmdt0(xf) = szRzazndt0(xf)) & isMonotone0(xf) & ! [X0, X1] : ((aElementOf0(X1, szDzozmdt0(xf)) & aElementOf0(X0, szDzozmdt0(xf))) => (sdtlseqdt0(X0, X1) => sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)))) & aFunction0(xf) & aCompleteLattice0(xU) & ! [X0] : ((aSubsetOf0(X0, xU) | (! [X1] : (aElementOf0(X1, X0) => aElementOf0(X1, xU)) & aSet0(X0))) => ? [X1] : (? [X2] : (aSupremumOfIn0(X2, X0, xU) & ! [X3] : ((aUpperBoundOfIn0(X3, X0, xU) | (! [X4] : (aElementOf0(X4, X0) => sdtlseqdt0(X4, X3)) & aElementOf0(X3, xU))) => sdtlseqdt0(X2, X3)) & aUpperBoundOfIn0(X2, X0, xU) & ! [X3] : (aElementOf0(X3, X0) => sdtlseqdt0(X3, X2)) & aElementOf0(X2, xU) & aElementOf0(X2, xU)) & aInfimumOfIn0(X1, X0, xU) & ! [X2] : ((aLowerBoundOfIn0(X2, X0, xU) | (! [X3] : (aElementOf0(X3, X0) => sdtlseqdt0(X2, X3)) & aElementOf0(X2, xU))) => sdtlseqdt0(X2, X1)) & aLowerBoundOfIn0(X1, X0, xU) & ! [X2] : (aElementOf0(X2, X0) => sdtlseqdt0(X1, X2)) & aElementOf0(X1, xU) & aElementOf0(X1, xU))) & aSet0(xU)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT387+4.p', m__1123)).
fof(f221, plain, (szDzozmdt0(xf) = szRzazndt0(xf)), inference(cnf_transformation, [], [f137])).
fof(f225, plain, ! [X0] : (aElementOf0(X0, szDzozmdt0(xf)) | ~ aElementOf0(X0, xS)), inference(cnf_transformation, [], [f69])).
fof(f3726, plain, (~ sdtlseqdt0(sK21, sK21) | ~ sdtlseqdt0(sK19(sK21), sK21) | aElementOf0(sK21, xP) | ~ aElementOf0(sK21, xU) | ~ spl23_5), inference(superposition, [], [f240, f3584])).
fof(f3584, plain, ((sK21 = sdtlpdtrp0(xf, sK21)) | ~ spl23_5), inference(resolution, [], [f299, f226])).
fof(f226, plain, ! [X0] : (~ aElementOf0(X0, xS) | (sdtlpdtrp0(xf, X0) = X0)), inference(cnf_transformation, [], [f69])).
fof(f240, plain, ! [X0] : (~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ sdtlseqdt0(sK19(X0), X0) | aElementOf0(X0, xP) | ~ aElementOf0(X0, xU)), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((aElementOf0(X0, xP) | (~ aUpperBoundOfIn0(X0, xT, xU) & (~ sdtlseqdt0(sK19(X0), X0) & aElementOf0(sK19(X0), xT))) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xU)) & ((aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (sdtlseqdt0(X2, X0) | ~ aElementOf0(X2, xT)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) | ~ aElementOf0(X0, xP))) & aSet0(xP)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK19])], [f72, f138])).
fof(f138, plain, ! [X0] : (? [X1] : (~ sdtlseqdt0(X1, X0) & aElementOf0(X1, xT)) => (~ sdtlseqdt0(sK19(X0), X0) & aElementOf0(sK19(X0), xT))), introduced(choice_axiom, [])).
fof(f72, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((aElementOf0(X0, xP) | (~ aUpperBoundOfIn0(X0, xT, xU) & ? [X1] : (~ sdtlseqdt0(X1, X0) & aElementOf0(X1, xT))) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xU)) & ((aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (sdtlseqdt0(X2, X0) | ~ aElementOf0(X2, xT)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) | ~ aElementOf0(X0, xP))) & aSet0(xP)), inference(flattening, [], [f71])).
fof(f71, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((aElementOf0(X0, xP) | ((~ aUpperBoundOfIn0(X0, xT, xU) & ? [X1] : (~ sdtlseqdt0(X1, X0) & aElementOf0(X1, xT))) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xU))) & ((aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (sdtlseqdt0(X2, X0) | ~ aElementOf0(X2, xT)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) | ~ aElementOf0(X0, xP))) & aSet0(xP)), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((((aUpperBoundOfIn0(X0, xT, xU) | ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0))) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) => aElementOf0(X0, xP)) & (aElementOf0(X0, xP) => (aUpperBoundOfIn0(X0, xT, xU) & ! [X2] : (aElementOf0(X2, xT) => sdtlseqdt0(X2, X0)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)))) & aSet0(xP)), inference(rectify, [], [f27])).
fof(f27, plain, ((xP = cS1241(xU, xf, xT)) & ! [X0] : ((((aUpperBoundOfIn0(X0, xT, xU) | ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0))) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)) => aElementOf0(X0, xP)) & (aElementOf0(X0, xP) => (aUpperBoundOfIn0(X0, xT, xU) & ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0)) & sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) & aElementOf0(X0, xU)))) & aSet0(xP)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT387+4.p', m__1244)).
fof(f3756, plain, (spl23_223 | spl23_225 | ~ spl23_224 | ~ spl23_5), inference(avatar_split_clause, [], [f3751, f297, f3747, f3753, f3743])).
fof(f3751, plain, (~ sdtlseqdt0(sK21, sK21) | aElementOf0(sK19(sK21), xT) | aElementOf0(sK21, xP) | ~ spl23_5), inference(subsumption_resolution, [], [f3725, f3585])).
fof(f3725, plain, (~ sdtlseqdt0(sK21, sK21) | aElementOf0(sK19(sK21), xT) | aElementOf0(sK21, xP) | ~ aElementOf0(sK21, xU) | ~ spl23_5), inference(superposition, [], [f239, f3584])).
fof(f239, plain, ! [X0] : (~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | aElementOf0(sK19(X0), xT) | aElementOf0(X0, xP) | ~ aElementOf0(X0, xU)), inference(cnf_transformation, [], [f139])).
fof(f3633, plain, (spl23_10 | ~ spl23_8 | ~ spl23_11), inference(avatar_split_clause, [], [f3617, f334, f313, f326])).
fof(f326, plain, (spl23_10 <=> sdtlseqdt0(sK22, xp)), introduced(avatar_definition, [new_symbols(naming, [spl23_10])])).
fof(f313, plain, (spl23_8 <=> (xp = sdtlpdtrp0(xf, xp))), introduced(avatar_definition, [new_symbols(naming, [spl23_8])])).
fof(f334, plain, (spl23_11 <=> aElementOf0(sK22, xT)), introduced(avatar_definition, [new_symbols(naming, [spl23_11])])).
fof(f3617, plain, (sdtlseqdt0(sK22, xp) | (~ spl23_8 | ~ spl23_11)), inference(resolution, [], [f3346, f336])).
fof(f336, plain, (aElementOf0(sK22, xT) | ~ spl23_11), inference(avatar_component_clause, [], [f334])).
fof(f3346, plain, (! [X0] : (~ aElementOf0(X0, xT) | sdtlseqdt0(X0, xp)) | ~ spl23_8), inference(backward_demodulation, [], [f253, f314])).
fof(f314, plain, ((xp = sdtlpdtrp0(xf, xp)) | ~ spl23_8), inference(avatar_component_clause, [], [f313])).
fof(f253, plain, ! [X0] : (sdtlseqdt0(X0, sdtlpdtrp0(xf, xp)) | ~ aElementOf0(X0, xT)), inference(cnf_transformation, [], [f74])).
fof(f74, plain, (aUpperBoundOfIn0(sdtlpdtrp0(xf, xp), xT, xU) & ! [X0] : (sdtlseqdt0(X0, sdtlpdtrp0(xf, xp)) | ~ aElementOf0(X0, xT)) & aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) & ! [X1] : (sdtlseqdt0(sdtlpdtrp0(xf, xp), X1) | ~ aElementOf0(X1, xP))), inference(ennf_transformation, [], [f39])).
fof(f39, plain, (aUpperBoundOfIn0(sdtlpdtrp0(xf, xp), xT, xU) & ! [X0] : (aElementOf0(X0, xT) => sdtlseqdt0(X0, sdtlpdtrp0(xf, xp))) & aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) & ! [X1] : (aElementOf0(X1, xP) => sdtlseqdt0(sdtlpdtrp0(xf, xp), X1))), inference(rectify, [], [f29])).
fof(f29, plain, (aUpperBoundOfIn0(sdtlpdtrp0(xf, xp), xT, xU) & ! [X0] : (aElementOf0(X0, xT) => sdtlseqdt0(X0, sdtlpdtrp0(xf, xp))) & aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU) & ! [X0] : (aElementOf0(X0, xP) => sdtlseqdt0(sdtlpdtrp0(xf, xp), X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT387+4.p', m__1299)).
fof(f3313, plain, (spl23_33 | ~ spl23_82), inference(avatar_split_clause, [], [f3304, f1299, f664])).
fof(f664, plain, (spl23_33 <=> sdtlseqdt0(xp, sdtlpdtrp0(xf, xp))), introduced(avatar_definition, [new_symbols(naming, [spl23_33])])).
fof(f1299, plain, (spl23_82 <=> aElementOf0(sdtlpdtrp0(xf, xp), xP)), introduced(avatar_definition, [new_symbols(naming, [spl23_82])])).
fof(f3304, plain, (sdtlseqdt0(xp, sdtlpdtrp0(xf, xp)) | ~ spl23_82), inference(resolution, [], [f1301, f245])).
fof(f1301, plain, (aElementOf0(sdtlpdtrp0(xf, xp), xP) | ~ spl23_82), inference(avatar_component_clause, [], [f1299])).
fof(f3312, plain, (~ spl23_33 | spl23_8 | ~ spl23_31), inference(avatar_split_clause, [], [f3311, f653, f313, f664])).
fof(f653, plain, (spl23_31 <=> aElement0(sdtlpdtrp0(xf, xp))), introduced(avatar_definition, [new_symbols(naming, [spl23_31])])).
fof(f3311, plain, ((xp = sdtlpdtrp0(xf, xp)) | ~ sdtlseqdt0(xp, sdtlpdtrp0(xf, xp)) | ~ spl23_31), inference(subsumption_resolution, [], [f3310, f349])).
fof(f349, plain, aElement0(xp), inference(subsumption_resolution, [], [f346, f213])).
fof(f213, plain, aSet0(xU), inference(cnf_transformation, [], [f137])).
fof(f346, plain, (aElement0(xp) | ~ aSet0(xU)), inference(resolution, [], [f148, f244])).
fof(f244, plain, aElementOf0(xp, xU), inference(cnf_transformation, [], [f141])).
fof(f3310, plain, ((xp = sdtlpdtrp0(xf, xp)) | ~ sdtlseqdt0(xp, sdtlpdtrp0(xf, xp)) | ~ aElement0(xp) | ~ spl23_31), inference(subsumption_resolution, [], [f643, f654])).
fof(f654, plain, (aElement0(sdtlpdtrp0(xf, xp)) | ~ spl23_31), inference(avatar_component_clause, [], [f653])).
fof(f643, plain, ((xp = sdtlpdtrp0(xf, xp)) | ~ sdtlseqdt0(xp, sdtlpdtrp0(xf, xp)) | ~ aElement0(sdtlpdtrp0(xf, xp)) | ~ aElement0(xp)), inference(resolution, [], [f156, f345])).
fof(f345, plain, sdtlseqdt0(sdtlpdtrp0(xf, xp), xp), inference(resolution, [], [f249, f252])).
fof(f252, plain, aLowerBoundOfIn0(sdtlpdtrp0(xf, xp), xP, xU), inference(cnf_transformation, [], [f74])).
fof(f249, plain, ! [X0] : (~ aLowerBoundOfIn0(X0, xP, xU) | sdtlseqdt0(X0, xp)), inference(cnf_transformation, [], [f141])).
fof(f156, plain, ! [X0, X1] : (~ sdtlseqdt0(X1, X0) | (X0 = X1) | ~ sdtlseqdt0(X0, X1) | ~ aElement0(X1) | ~ aElement0(X0)), inference(cnf_transformation, [], [f47])).
fof(f47, plain, ! [X0, X1] : ((X0 = X1) | ~ sdtlseqdt0(X1, X0) | ~ sdtlseqdt0(X0, X1) | ~ aElement0(X1) | ~ aElement0(X0)), inference(flattening, [], [f46])).
fof(f46, plain, ! [X0, X1] : (((X0 = X1) | (~ sdtlseqdt0(X1, X0) | ~ sdtlseqdt0(X0, X1))) | (~ aElement0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1] : ((aElement0(X1) & aElement0(X0)) => ((sdtlseqdt0(X1, X0) & sdtlseqdt0(X0, X1)) => (X0 = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT387+4.p', mASymm)).
fof(f3286, plain, spl23_82, inference(avatar_split_clause, [], [f3283, f1299])).
fof(f3283, plain, aElementOf0(sdtlpdtrp0(xf, xp), xP), inference(subsumption_resolution, [], [f3282, f244])).
fof(f3282, plain, (~ aElementOf0(xp, xU) | aElementOf0(sdtlpdtrp0(xf, xp), xP)), inference(subsumption_resolution, [], [f3273, f345])).
fof(f3273, plain, (~ sdtlseqdt0(sdtlpdtrp0(xf, xp), xp) | ~ aElementOf0(xp, xU) | aElementOf0(sdtlpdtrp0(xf, xp), xP)), inference(resolution, [], [f877, f254])).
fof(f254, plain, aUpperBoundOfIn0(sdtlpdtrp0(xf, xp), xT, xU), inference(cnf_transformation, [], [f74])).
fof(f877, plain, ! [X0] : (~ aUpperBoundOfIn0(sdtlpdtrp0(xf, X0), xT, xU) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aElementOf0(X0, xU) | aElementOf0(sdtlpdtrp0(xf, X0), xP)), inference(subsumption_resolution, [], [f876, f627])).
fof(f627, plain, ! [X0] : (aElementOf0(sdtlpdtrp0(xf, X0), xU) | ~ aElementOf0(X0, xU)), inference(forward_demodulation, [], [f626, f273])).
fof(f626, plain, ! [X0] : (aElementOf0(sdtlpdtrp0(xf, X0), xU) | ~ aElementOf0(X0, szDzozmdt0(xf))), inference(subsumption_resolution, [], [f623, f218])).
fof(f218, plain, aFunction0(xf), inference(cnf_transformation, [], [f137])).
fof(f623, plain, ! [X0] : (aElementOf0(sdtlpdtrp0(xf, X0), xU) | ~ aElementOf0(X0, szDzozmdt0(xf)) | ~ aFunction0(xf)), inference(superposition, [], [f186, f222])).
fof(f186, plain, ! [X0, X1] : (aElementOf0(sdtlpdtrp0(X0, X1), szRzazndt0(X0)) | ~ aElementOf0(X1, szDzozmdt0(X0)) | ~ aFunction0(X0)), inference(cnf_transformation, [], [f63])).
fof(f63, plain, ! [X0] : (! [X1] : (aElementOf0(sdtlpdtrp0(X0, X1), szRzazndt0(X0)) | ~ aElementOf0(X1, szDzozmdt0(X0))) | ~ aFunction0(X0)), inference(ennf_transformation, [], [f21])).
fof(f21, plain, ! [X0] : (aFunction0(X0) => ! [X1] : (aElementOf0(X1, szDzozmdt0(X0)) => aElementOf0(sdtlpdtrp0(X0, X1), szRzazndt0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT387+4.p', mImgSort)).
fof(f876, plain, ! [X0] : (~ aElementOf0(X0, xU) | ~ aElementOf0(sdtlpdtrp0(xf, X0), xU) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aUpperBoundOfIn0(sdtlpdtrp0(xf, X0), xT, xU) | aElementOf0(sdtlpdtrp0(xf, X0), xP)), inference(duplicate_literal_removal, [], [f869])).
fof(f869, plain, ! [X0] : (~ aElementOf0(X0, xU) | ~ aElementOf0(sdtlpdtrp0(xf, X0), xU) | ~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aUpperBoundOfIn0(sdtlpdtrp0(xf, X0), xT, xU) | aElementOf0(sdtlpdtrp0(xf, X0), xP) | ~ aElementOf0(sdtlpdtrp0(xf, X0), xU)), inference(resolution, [], [f275, f241])).
fof(f241, plain, ! [X0] : (~ sdtlseqdt0(sdtlpdtrp0(xf, X0), X0) | ~ aUpperBoundOfIn0(X0, xT, xU) | aElementOf0(X0, xP) | ~ aElementOf0(X0, xU)), inference(cnf_transformation, [], [f139])).
fof(f275, plain, ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ aElementOf0(X1, xU) | ~ aElementOf0(X0, xU) | ~ sdtlseqdt0(X0, X1)), inference(forward_demodulation, [], [f274, f273])).
fof(f274, plain, ! [X0, X1] : (~ aElementOf0(X1, xU) | sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X0, szDzozmdt0(xf))), inference(forward_demodulation, [], [f219, f273])).
fof(f219, plain, ! [X0, X1] : (sdtlseqdt0(sdtlpdtrp0(xf, X0), sdtlpdtrp0(xf, X1)) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szDzozmdt0(xf)) | ~ aElementOf0(X0, szDzozmdt0(xf))), inference(cnf_transformation, [], [f137])).
fof(f1323, plain, (spl23_31 | ~ spl23_29), inference(avatar_split_clause, [], [f1322, f585, f653])).
fof(f585, plain, (spl23_29 <=> aElementOf0(sdtlpdtrp0(xf, xp), xU)), introduced(avatar_definition, [new_symbols(naming, [spl23_29])])).
fof(f1322, plain, (aElement0(sdtlpdtrp0(xf, xp)) | ~ spl23_29), inference(subsumption_resolution, [], [f1318, f213])).
fof(f1318, plain, (aElement0(sdtlpdtrp0(xf, xp)) | ~ aSet0(xU) | ~ spl23_29), inference(resolution, [], [f587, f148])).
fof(f587, plain, (aElementOf0(sdtlpdtrp0(xf, xp), xU) | ~ spl23_29), inference(avatar_component_clause, [], [f585])).
fof(f1287, plain, spl23_29, inference(avatar_contradiction_clause, [], [f1286])).
fof(f1286, plain, ($false | spl23_29), inference(subsumption_resolution, [], [f1282, f244])).
fof(f1282, plain, (~ aElementOf0(xp, xU) | spl23_29), inference(resolution, [], [f627, f586])).
fof(f586, plain, (~ aElementOf0(sdtlpdtrp0(xf, xp), xU) | spl23_29), inference(avatar_component_clause, [], [f585])).
fof(f340, plain, (~ spl23_8 | spl23_11 | spl23_1), inference(avatar_split_clause, [], [f339, f279, f334, f313])).
fof(f279, plain, (spl23_1 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl23_1])])).
fof(f339, plain, (sP3 | aElementOf0(sK22, xT) | ~ (xp = sdtlpdtrp0(xf, xp))), inference(subsumption_resolution, [], [f338, f244])).
fof(f338, plain, (~ aElementOf0(xp, xU) | sP3 | aElementOf0(sK22, xT) | ~ (xp = sdtlpdtrp0(xf, xp))), inference(forward_demodulation, [], [f259, f273])).
fof(f259, plain, (sP3 | aElementOf0(sK22, xT) | ~ (xp = sdtlpdtrp0(xf, xp)) | ~ aElementOf0(xp, szDzozmdt0(xf))), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ((~ aSupremumOfIn0(xp, xT, xS) & (sP3 | (~ aUpperBoundOfIn0(xp, xT, xS) & (~ sdtlseqdt0(sK22, xp) & aElementOf0(sK22, xT))))) | (~ aFixedPointOf0(xp, xf) & (~ (xp = sdtlpdtrp0(xf, xp)) | ~ aElementOf0(xp, szDzozmdt0(xf))))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK22])], [f145, f146])).
fof(f146, plain, (? [X0] : (~ sdtlseqdt0(X0, xp) & aElementOf0(X0, xT)) => (~ sdtlseqdt0(sK22, xp) & aElementOf0(sK22, xT))), introduced(choice_axiom, [])).
fof(f145, plain, ((~ aSupremumOfIn0(xp, xT, xS) & (sP3 | (~ aUpperBoundOfIn0(xp, xT, xS) & ? [X0] : (~ sdtlseqdt0(X0, xp) & aElementOf0(X0, xT))))) | (~ aFixedPointOf0(xp, xf) & (~ (xp = sdtlpdtrp0(xf, xp)) | ~ aElementOf0(xp, szDzozmdt0(xf))))), inference(rectify, [], [f82])).
fof(f82, plain, ((~ aSupremumOfIn0(xp, xT, xS) & (sP3 | (~ aUpperBoundOfIn0(xp, xT, xS) & ? [X2] : (~ sdtlseqdt0(X2, xp) & aElementOf0(X2, xT))))) | (~ aFixedPointOf0(xp, xf) & (~ (xp = sdtlpdtrp0(xf, xp)) | ~ aElementOf0(xp, szDzozmdt0(xf))))), inference(definition_folding, [], [f76, e81])).
fof(f81, plain, (? [X0] : (~ sdtlseqdt0(xp, X0) & aUpperBoundOfIn0(X0, xT, xS) & ! [X1] : (sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, xT)) & aElementOf0(X0, xS)) | ~ sP3), inference(usedef, [], [e81])).
fof(e81, plain, (sP3 <=> ? [X0] : (~ sdtlseqdt0(xp, X0) & aUpperBoundOfIn0(X0, xT, xS) & ! [X1] : (sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, xT)) & aElementOf0(X0, xS))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f76, plain, ((~ aSupremumOfIn0(xp, xT, xS) & (? [X0] : (~ sdtlseqdt0(xp, X0) & aUpperBoundOfIn0(X0, xT, xS) & ! [X1] : (sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, xT)) & aElementOf0(X0, xS)) | (~ aUpperBoundOfIn0(xp, xT, xS) & ? [X2] : (~ sdtlseqdt0(X2, xp) & aElementOf0(X2, xT))))) | (~ aFixedPointOf0(xp, xf) & (~ (xp = sdtlpdtrp0(xf, xp)) | ~ aElementOf0(xp, szDzozmdt0(xf))))), inference(flattening, [], [f75])).
fof(f75, plain, ((~ aSupremumOfIn0(xp, xT, xS) & (? [X0] : (~ sdtlseqdt0(xp, X0) & (aUpperBoundOfIn0(X0, xT, xS) & ! [X1] : (sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, xT)) & aElementOf0(X0, xS))) | (~ aUpperBoundOfIn0(xp, xT, xS) & ? [X2] : (~ sdtlseqdt0(X2, xp) & aElementOf0(X2, xT))))) | (~ aFixedPointOf0(xp, xf) & (~ (xp = sdtlpdtrp0(xf, xp)) | ~ aElementOf0(xp, szDzozmdt0(xf))))), inference(ennf_transformation, [], [f40])).
fof(f40, plain, ~ ((aSupremumOfIn0(xp, xT, xS) | (! [X0] : ((aUpperBoundOfIn0(X0, xT, xS) & ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0)) & aElementOf0(X0, xS)) => sdtlseqdt0(xp, X0)) & (aUpperBoundOfIn0(xp, xT, xS) | ! [X2] : (aElementOf0(X2, xT) => sdtlseqdt0(X2, xp))))) & (aFixedPointOf0(xp, xf) | ((xp = sdtlpdtrp0(xf, xp)) & aElementOf0(xp, szDzozmdt0(xf))))), inference(rectify, [], [f31])).
fof(f31, plain, ~ ((aSupremumOfIn0(xp, xT, xS) | (! [X0] : ((aUpperBoundOfIn0(X0, xT, xS) & ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0)) & aElementOf0(X0, xS)) => sdtlseqdt0(xp, X0)) & (aUpperBoundOfIn0(xp, xT, xS) | ! [X0] : (aElementOf0(X0, xT) => sdtlseqdt0(X0, xp))))) & (aFixedPointOf0(xp, xf) | ((xp = sdtlpdtrp0(xf, xp)) & aElementOf0(xp, szDzozmdt0(xf))))), inference(negated_conjecture, [], [f30])).
fof(f30, plain, ~ ((aSupremumOfIn0(xp, xT, xS) | (! [X0] : ((aUpperBoundOfIn0(X0, xT, xS) & ! [X1] : (aElementOf0(X1, xT) => sdtlseqdt0(X1, X0)) & aElementOf0(X0, xS)) => sdtlseqdt0(xp, X0)) & (aUpperBoundOfIn0(xp, xT, xS) | ! [X0] : (aElementOf0(X0, xT) => sdtlseqdt0(X0, xp))))) & (aFixedPointOf0(xp, xf) | ((xp = sdtlpdtrp0(xf, xp)) & aElementOf0(xp, szDzozmdt0(xf))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT387+4.p', m__)).
fof(f332, plain, (~ spl23_8 | ~ spl23_10 | spl23_1), inference(avatar_split_clause, [], [f331, f279, f326, f313])).
fof(f331, plain, (sP3 | ~ sdtlseqdt0(sK22, xp) | ~ (xp = sdtlpdtrp0(xf, xp))), inference(subsumption_resolution, [], [f330, f244])).
fof(f330, plain, (~ aElementOf0(xp, xU) | sP3 | ~ sdtlseqdt0(sK22, xp) | ~ (xp = sdtlpdtrp0(xf, xp))), inference(forward_demodulation, [], [f261, f273])).
fof(f261, plain, (sP3 | ~ sdtlseqdt0(sK22, xp) | ~ (xp = sdtlpdtrp0(xf, xp)) | ~ aElementOf0(xp, szDzozmdt0(xf))), inference(cnf_transformation, [], [f147])).
fof(f300, plain, (~ spl23_1 | spl23_5), inference(avatar_split_clause, [], [f255, f297, f279])).
fof(f255, plain, (aElementOf0(sK21, xS) | ~ sP3), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ((~ sdtlseqdt0(xp, sK21) & aUpperBoundOfIn0(sK21, xT, xS) & ! [X1] : (sdtlseqdt0(X1, sK21) | ~ aElementOf0(X1, xT)) & aElementOf0(sK21, xS)) | ~ sP3), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21])], [f142, f143])).
fof(f143, plain, (? [X0] : (~ sdtlseqdt0(xp, X0) & aUpperBoundOfIn0(X0, xT, xS) & ! [X1] : (sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, xT)) & aElementOf0(X0, xS)) => (~ sdtlseqdt0(xp, sK21) & aUpperBoundOfIn0(sK21, xT, xS) & ! [X1] : (sdtlseqdt0(X1, sK21) | ~ aElementOf0(X1, xT)) & aElementOf0(sK21, xS))), introduced(choice_axiom, [])).
fof(f142, plain, (? [X0] : (~ sdtlseqdt0(xp, X0) & aUpperBoundOfIn0(X0, xT, xS) & ! [X1] : (sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, xT)) & aElementOf0(X0, xS)) | ~ sP3), inference(nnf_transformation, [], [f81])).
fof(f295, plain, (~ spl23_1 | spl23_4), inference(avatar_split_clause, [], [f256, f293, f279])).
fof(f256, plain, ! [X1] : (sdtlseqdt0(X1, sK21) | ~ aElementOf0(X1, xT) | ~ sP3), inference(cnf_transformation, [], [f144])).
fof(f286, plain, (~ spl23_1 | ~ spl23_2), inference(avatar_split_clause, [], [f258, f283, f279])).
fof(f258, plain, (~ sdtlseqdt0(xp, sK21) | ~ sP3), inference(cnf_transformation, [], [f144])).