fof(f45116, plain, $false, inference(avatar_sat_refutation, [], [f506, f2198, f2376, f2487, f45036, f45114])).
fof(f45114, plain, (spl33_2 | ~ spl33_72 | ~ spl33_100), inference(avatar_contradiction_clause, [], [f45113])).
fof(f45113, plain, ($false | (spl33_2 | ~ spl33_72 | ~ spl33_100)), inference(subsumption_resolution, [], [f45112, f45034])).
fof(f45034, plain, ((k23_filter_2(sK4, k17_filter_2(sK4)) = g3_lattices(k17_filter_2(sK4), u2_lattices(sK4), u1_lattices(sK4))) | (~ spl33_72 | ~ spl33_100)), inference(forward_demodulation, [], [f45033, f29714])).
fof(f29714, plain, ((u2_lattices(sK4) = sK5(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4)) | (~ spl33_72 | ~ spl33_100)), inference(forward_demodulation, [], [f29713, f4094])).
fof(f4094, plain, ((u2_lattices(sK4) = k1_realset1(u2_lattices(sK4), k17_filter_2(sK4))) | (~ spl33_72 | ~ spl33_100)), inference(forward_demodulation, [], [f4089, f3197])).
fof(f3197, plain, (u2_lattices(sK4) = k7_relat_1(u2_lattices(sK4), k1_relat_1(u2_lattices(sK4)))), inference(resolution, [], [f3193, f828])).
fof(f828, plain, ! [X0] : (~ v1_relat_1(X0) | (k7_relat_1(X0, k1_relat_1(X0)) = X0)), inference(resolution, [], [f486, f475])).
fof(f475, plain, ! [X0] : r1_tarski(X0, X0), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ! [X0] : r1_tarski(X0, X0), inference(rectify, [], [f90])).
fof(f90, plain, ! [X0, X1] : r1_tarski(X0, X0), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', reflexivity_r1_tarski)).
fof(f486, plain, ! [X0, X1] : (~ r1_tarski(k1_relat_1(X1), X0) | (k7_relat_1(X1, X0) = X1) | ~ v1_relat_1(X1)), inference(cnf_transformation, [], [f225])).
fof(f225, plain, ! [X0, X1] : ((k7_relat_1(X1, X0) = X1) | ~ r1_tarski(k1_relat_1(X1), X0) | ~ v1_relat_1(X1)), inference(flattening, [], [f224])).
fof(f224, plain, ! [X0, X1] : (((k7_relat_1(X1, X0) = X1) | ~ r1_tarski(k1_relat_1(X1), X0)) | ~ v1_relat_1(X1)), inference(ennf_transformation, [], [f100])).
fof(f100, plain, ! [X0, X1] : (v1_relat_1(X1) => (r1_tarski(k1_relat_1(X1), X0) => (k7_relat_1(X1, X0) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', t97_relat_1)).
fof(f3193, plain, v1_relat_1(u2_lattices(sK4)), inference(resolution, [], [f887, f684])).
fof(f684, plain, ! [X2, X0, X1] : (~ m2_relset_1(X0, X1, X2) | v1_relat_1(X0)), inference(resolution, [], [f365, f315])).
fof(f315, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, k1_zfmisc_1(k2_zfmisc_1(X0, X1))) | v1_relat_1(X2)), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ! [X0, X1, X2] : (v1_relat_1(X2) | ~ m1_subset_1(X2, k1_zfmisc_1(k2_zfmisc_1(X0, X1)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ! [X0, X1, X2] : (m1_subset_1(X2, k1_zfmisc_1(k2_zfmisc_1(X0, X1))) => v1_relat_1(X2)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', cc1_relset_1)).
fof(f365, plain, ! [X2, X0, X1] : (m1_subset_1(X2, k1_zfmisc_1(k2_zfmisc_1(X0, X1))) | ~ m2_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ! [X0, X1, X2] : (m1_subset_1(X2, k1_zfmisc_1(k2_zfmisc_1(X0, X1))) | ~ m2_relset_1(X2, X0, X1)), inference(ennf_transformation, [], [f39])).
fof(f39, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) => m1_subset_1(X2, k1_zfmisc_1(k2_zfmisc_1(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', dt_m2_relset_1)).
fof(f887, plain, m2_relset_1(u2_lattices(sK4), k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4)), inference(subsumption_resolution, [], [f885, f519])).
fof(f519, plain, l2_lattices(sK4), inference(resolution, [], [f354, f303])).
fof(f303, plain, l3_lattices(sK4), inference(cnf_transformation, [], [f234])).
fof(f234, plain, ((~ (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k1_filter_2(sK4))) | ~ (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k17_filter_2(sK4)))) & l3_lattices(sK4) & v10_lattices(sK4) & ~ v3_struct_0(sK4)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4])], [f113, f233])).
fof(f233, plain, (? [X0] : ((~ (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = k23_filter_2(X0, k1_filter_2(X0))) | ~ (k23_filter_2(X0, k17_filter_2(X0)) = g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)))) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ((~ (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k1_filter_2(sK4))) | ~ (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k17_filter_2(sK4)))) & l3_lattices(sK4) & v10_lattices(sK4) & ~ v3_struct_0(sK4))), introduced(choice_axiom, [])).
fof(f113, plain, ? [X0] : ((~ (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = k23_filter_2(X0, k1_filter_2(X0))) | ~ (k23_filter_2(X0, k17_filter_2(X0)) = g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)))) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f112])).
fof(f112, plain, ? [X0] : ((~ (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = k23_filter_2(X0, k1_filter_2(X0))) | ~ (k23_filter_2(X0, k17_filter_2(X0)) = g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)))) & (l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = k23_filter_2(X0, k1_filter_2(X0))) & (k23_filter_2(X0, k17_filter_2(X0)) = g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = k23_filter_2(X0, k1_filter_2(X0))) & (k23_filter_2(X0, k17_filter_2(X0)) = g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', t72_filter_2)).
fof(f354, plain, ! [X0] : (~ l3_lattices(X0) | l2_lattices(X0)), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ! [X0] : ((l2_lattices(X0) & l1_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', dt_l3_lattices)).
fof(f885, plain, (m2_relset_1(u2_lattices(sK4), k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4)) | ~ l2_lattices(sK4)), inference(superposition, [], [f371, f755])).
fof(f755, plain, (u1_struct_0(sK4) = k17_filter_2(sK4)), inference(subsumption_resolution, [], [f754, f301])).
fof(f301, plain, ~ v3_struct_0(sK4), inference(cnf_transformation, [], [f234])).
fof(f754, plain, ((u1_struct_0(sK4) = k17_filter_2(sK4)) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f749, f302])).
fof(f302, plain, v10_lattices(sK4), inference(cnf_transformation, [], [f234])).
fof(f749, plain, ((u1_struct_0(sK4) = k17_filter_2(sK4)) | ~ v10_lattices(sK4) | v3_struct_0(sK4)), inference(resolution, [], [f341, f303])).
fof(f341, plain, ! [X0] : (~ l3_lattices(X0) | (k17_filter_2(X0) = u1_struct_0(X0)) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ! [X0] : ((k17_filter_2(X0) = u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f132])).
fof(f132, plain, ! [X0] : ((k17_filter_2(X0) = u1_struct_0(X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k17_filter_2(X0) = u1_struct_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', d8_filter_2)).
fof(f371, plain, ! [X0] : (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ! [X0] : ((m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0))) | ~ l2_lattices(X0)), inference(ennf_transformation, [], [f42])).
fof(f42, plain, ! [X0] : (l2_lattices(X0) => (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', dt_u2_lattices)).
fof(f4089, plain, ((k7_relat_1(u2_lattices(sK4), k1_relat_1(u2_lattices(sK4))) = k1_realset1(u2_lattices(sK4), k17_filter_2(sK4))) | (~ spl33_72 | ~ spl33_100)), inference(superposition, [], [f2024, f3283])).
fof(f3283, plain, ((k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)) = k1_relat_1(u2_lattices(sK4))) | (~ spl33_72 | ~ spl33_100)), inference(backward_demodulation, [], [f3225, f3281])).
fof(f3281, plain, ((k1_relat_1(u1_lattices(sK4)) = k1_relat_1(u2_lattices(sK4))) | (~ spl33_72 | ~ spl33_100)), inference(backward_demodulation, [], [f3232, f3279])).
fof(f3279, plain, ((k1_relat_1(u2_lattices(sK4)) = k4_relset_1(k1_relat_1(u1_lattices(sK4)), k17_filter_2(sK4), u2_lattices(sK4))) | ~ spl33_72), inference(resolution, [], [f3231, f466])).
fof(f466, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, X0, X1) | (k4_relset_1(X0, X1, X2) = k1_relat_1(X2))), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ! [X0, X1, X2] : ((k4_relset_1(X0, X1, X2) = k1_relat_1(X2)) | ~ m1_relset_1(X2, X0, X1)), inference(ennf_transformation, [], [f84])).
fof(f84, plain, ! [X0, X1, X2] : (m1_relset_1(X2, X0, X1) => (k4_relset_1(X0, X1, X2) = k1_relat_1(X2))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', redefinition_k4_relset_1)).
fof(f3231, plain, (m1_relset_1(u2_lattices(sK4), k1_relat_1(u1_lattices(sK4)), k17_filter_2(sK4)) | ~ spl33_72), inference(backward_demodulation, [], [f1711, f3225])).
fof(f1711, plain, m1_relset_1(u2_lattices(sK4), k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4)), inference(subsumption_resolution, [], [f1703, f519])).
fof(f1703, plain, (m1_relset_1(u2_lattices(sK4), k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4)) | ~ l2_lattices(sK4)), inference(superposition, [], [f884, f755])).
fof(f884, plain, ! [X1] : (m1_relset_1(u2_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l2_lattices(X1)), inference(resolution, [], [f371, f470])).
fof(f470, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f298])).
fof(f298, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f87])).
fof(f87, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', redefinition_m2_relset_1)).
fof(f3232, plain, ((k1_relat_1(u1_lattices(sK4)) = k4_relset_1(k1_relat_1(u1_lattices(sK4)), k17_filter_2(sK4), u2_lattices(sK4))) | (~ spl33_72 | ~ spl33_100)), inference(backward_demodulation, [], [f2375, f3225])).
fof(f2375, plain, ((k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)) = k4_relset_1(k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4), u2_lattices(sK4))) | ~ spl33_100), inference(avatar_component_clause, [], [f2373])).
fof(f2373, plain, (spl33_100 <=> (k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)) = k4_relset_1(k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4), u2_lattices(sK4)))), introduced(avatar_definition, [new_symbols(naming, [spl33_100])])).
fof(f3225, plain, ((k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)) = k1_relat_1(u1_lattices(sK4))) | ~ spl33_72), inference(backward_demodulation, [], [f2193, f3203])).
fof(f3203, plain, (k4_relset_1(k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4), u1_lattices(sK4)) = k1_relat_1(u1_lattices(sK4))), inference(resolution, [], [f1691, f466])).
fof(f1691, plain, m1_relset_1(u1_lattices(sK4), k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4)), inference(subsumption_resolution, [], [f1683, f513])).
fof(f513, plain, l1_lattices(sK4), inference(resolution, [], [f353, f303])).
fof(f353, plain, ! [X0] : (~ l3_lattices(X0) | l1_lattices(X0)), inference(cnf_transformation, [], [f149])).
fof(f1683, plain, (m1_relset_1(u1_lattices(sK4), k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4)) | ~ l1_lattices(sK4)), inference(superposition, [], [f872, f755])).
fof(f872, plain, ! [X1] : (m1_relset_1(u1_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l1_lattices(X1)), inference(resolution, [], [f368, f470])).
fof(f368, plain, ! [X0] : (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ! [X0] : ((m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0))) | ~ l1_lattices(X0)), inference(ennf_transformation, [], [f40])).
fof(f40, plain, ! [X0] : (l1_lattices(X0) => (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', dt_u1_lattices)).
fof(f2193, plain, ((k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)) = k4_relset_1(k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4), u1_lattices(sK4))) | ~ spl33_72), inference(avatar_component_clause, [], [f2191])).
fof(f2191, plain, (spl33_72 <=> (k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)) = k4_relset_1(k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4), u1_lattices(sK4)))), introduced(avatar_definition, [new_symbols(naming, [spl33_72])])).
fof(f2024, plain, ! [X0] : (k1_realset1(u2_lattices(sK4), X0) = k7_relat_1(u2_lattices(sK4), k2_zfmisc_1(X0, X0))), inference(subsumption_resolution, [], [f2023, f301])).
fof(f2023, plain, ! [X0] : ((k1_realset1(u2_lattices(sK4), X0) = k7_relat_1(u2_lattices(sK4), k2_zfmisc_1(X0, X0))) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f2017, f599])).
fof(f599, plain, v5_lattices(sK4), inference(subsumption_resolution, [], [f598, f301])).
fof(f598, plain, (v3_struct_0(sK4) | v5_lattices(sK4)), inference(subsumption_resolution, [], [f593, f302])).
fof(f593, plain, (~ v10_lattices(sK4) | v3_struct_0(sK4) | v5_lattices(sK4)), inference(resolution, [], [f310, f303])).
fof(f310, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v5_lattices(X0)), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ! [X0] : ((v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f118])).
fof(f118, plain, ! [X0] : (((v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)) | (~ v10_lattices(X0) | v3_struct_0(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ! [X0] : (l3_lattices(X0) => ((v10_lattices(X0) & ~ v3_struct_0(X0)) => (v9_lattices(X0) & v8_lattices(X0) & v7_lattices(X0) & v6_lattices(X0) & v5_lattices(X0) & v4_lattices(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', cc1_lattices)).
fof(f2017, plain, ! [X0] : ((k1_realset1(u2_lattices(sK4), X0) = k7_relat_1(u2_lattices(sK4), k2_zfmisc_1(X0, X0))) | ~ v5_lattices(sK4) | v3_struct_0(sK4)), inference(resolution, [], [f736, f519])).
fof(f736, plain, ! [X0, X1] : (~ l2_lattices(X0) | (k1_realset1(u2_lattices(X0), X1) = k7_relat_1(u2_lattices(X0), k2_zfmisc_1(X1, X1))) | ~ v5_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f340, f393])).
fof(f393, plain, ! [X0] : (v1_relat_1(u2_lattices(X0)) | ~ l2_lattices(X0) | ~ v5_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ! [X0] : ((v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)) & v1_relat_1(u2_lattices(X0))) | ~ l2_lattices(X0) | ~ v5_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f180])).
fof(f180, plain, ! [X0] : ((v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)) & v1_relat_1(u2_lattices(X0))) | (~ l2_lattices(X0) | ~ v5_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f107])).
fof(f107, plain, ! [X0] : ((l2_lattices(X0) & v5_lattices(X0) & ~ v3_struct_0(X0)) => (v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)) & v1_relat_1(u2_lattices(X0)))), inference(pure_predicate_removal, [], [f102])).
fof(f102, plain, ! [X0] : ((l2_lattices(X0) & v5_lattices(X0) & ~ v3_struct_0(X0)) => (v1_partfun1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)) & v1_relat_1(u2_lattices(X0)))), inference(pure_predicate_removal, [], [f61])).
fof(f61, plain, ! [X0] : ((l2_lattices(X0) & v5_lattices(X0) & ~ v3_struct_0(X0)) => (v1_partfun1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v2_binop_1(u2_lattices(X0), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)) & v1_relat_1(u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', fc3_lattice2)).
fof(f340, plain, ! [X0, X1] : (~ v1_relat_1(X0) | (k1_realset1(X0, X1) = k7_relat_1(X0, k2_zfmisc_1(X1, X1)))), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ! [X0] : (! [X1] : (k1_realset1(X0, X1) = k7_relat_1(X0, k2_zfmisc_1(X1, X1))) | ~ v1_relat_1(X0)), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0] : (v1_relat_1(X0) => ! [X1] : (k1_realset1(X0, X1) = k7_relat_1(X0, k2_zfmisc_1(X1, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', d3_realset1)).
fof(f29713, plain, (k1_realset1(u2_lattices(sK4), k17_filter_2(sK4)) = sK5(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4)), inference(subsumption_resolution, [], [f29712, f301])).
fof(f29712, plain, (v3_struct_0(sK4) | (k1_realset1(u2_lattices(sK4), k17_filter_2(sK4)) = sK5(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4))), inference(subsumption_resolution, [], [f29699, f302])).
fof(f29699, plain, (~ v10_lattices(sK4) | v3_struct_0(sK4) | (k1_realset1(u2_lattices(sK4), k17_filter_2(sK4)) = sK5(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4))), inference(resolution, [], [f6331, f303])).
fof(f6331, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | (k1_realset1(u2_lattices(X0), k17_filter_2(X0)) = sK5(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0))), inference(subsumption_resolution, [], [f6330, f788])).
fof(f788, plain, ! [X0] : (~ v1_xboole_0(k17_filter_2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(duplicate_literal_removal, [], [f785])).
fof(f785, plain, ! [X0] : (~ v1_xboole_0(k17_filter_2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f359, f344])).
fof(f344, plain, ! [X0] : (m2_filter_2(k17_filter_2(X0), X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f137])).
fof(f137, plain, ! [X0] : (m2_filter_2(k17_filter_2(X0), X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f136])).
fof(f136, plain, ! [X0] : (m2_filter_2(k17_filter_2(X0), X0) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m2_filter_2(k17_filter_2(X0), X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', dt_k17_filter_2)).
fof(f359, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | ~ v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f154])).
fof(f154, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', dt_m2_filter_2)).
fof(f6330, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v1_xboole_0(k17_filter_2(X0)) | (k1_realset1(u2_lattices(X0), k17_filter_2(X0)) = sK5(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0))), inference(duplicate_literal_removal, [], [f6312])).
fof(f6312, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v1_xboole_0(k17_filter_2(X0)) | (k1_realset1(u2_lattices(X0), k17_filter_2(X0)) = sK5(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f2167, f859])).
fof(f859, plain, ! [X0] : (m2_lattice4(k17_filter_2(X0), X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(duplicate_literal_removal, [], [f856])).
fof(f856, plain, ! [X0] : (m2_lattice4(k17_filter_2(X0), X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f360, f344])).
fof(f360, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f155])).
fof(f2167, plain, ! [X4, X5] : (~ m2_lattice4(X4, X5) | ~ l3_lattices(X5) | ~ v10_lattices(X5) | v3_struct_0(X5) | v1_xboole_0(X4) | (k1_realset1(u2_lattices(X5), X4) = sK5(k23_filter_2(X5, X4), X4, X5))), inference(resolution, [], [f1103, f328])).
fof(f328, plain, ! [X2, X0, X1] : (~ sP0(X0, X1, X2) | (k1_realset1(u2_lattices(X2), X1) = sK5(X0, X1, X2))), inference(cnf_transformation, [], [f240])).
fof(f240, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | ! [X3] : (! [X4] : (~ (g3_lattices(X1, X3, X4) = X0) | ~ (k1_realset1(u1_lattices(X2), X1) = X4) | ~ (k1_realset1(u2_lattices(X2), X1) = X3) | ~ m2_relset_1(X4, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_2(X4, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_1(X4)) | ~ m2_relset_1(X3, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_2(X3, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_1(X3))) & ((((g3_lattices(X1, sK5(X0, X1, X2), sK6(X0, X1, X2)) = X0) & (k1_realset1(u1_lattices(X2), X1) = sK6(X0, X1, X2)) & (k1_realset1(u2_lattices(X2), X1) = sK5(X0, X1, X2)) & m2_relset_1(sK6(X0, X1, X2), k2_zfmisc_1(X1, X1), X1) & v1_funct_2(sK6(X0, X1, X2), k2_zfmisc_1(X1, X1), X1) & v1_funct_1(sK6(X0, X1, X2))) & m2_relset_1(sK5(X0, X1, X2), k2_zfmisc_1(X1, X1), X1) & v1_funct_2(sK5(X0, X1, X2), k2_zfmisc_1(X1, X1), X1) & v1_funct_1(sK5(X0, X1, X2))) | ~ sP0(X0, X1, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5, sK6])], [f237, f239, f238])).
fof(f238, plain, ! [X2, X1, X0] : (? [X5] : (? [X6] : ((g3_lattices(X1, X5, X6) = X0) & (k1_realset1(u1_lattices(X2), X1) = X6) & (k1_realset1(u2_lattices(X2), X1) = X5) & m2_relset_1(X6, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X6, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X6)) & m2_relset_1(X5, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X5, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X5)) => (? [X6] : ((g3_lattices(X1, sK5(X0, X1, X2), X6) = X0) & (k1_realset1(u1_lattices(X2), X1) = X6) & (k1_realset1(u2_lattices(X2), X1) = sK5(X0, X1, X2)) & m2_relset_1(X6, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X6, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X6)) & m2_relset_1(sK5(X0, X1, X2), k2_zfmisc_1(X1, X1), X1) & v1_funct_2(sK5(X0, X1, X2), k2_zfmisc_1(X1, X1), X1) & v1_funct_1(sK5(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f239, plain, ! [X2, X1, X0] : (? [X6] : ((g3_lattices(X1, sK5(X0, X1, X2), X6) = X0) & (k1_realset1(u1_lattices(X2), X1) = X6) & (k1_realset1(u2_lattices(X2), X1) = sK5(X0, X1, X2)) & m2_relset_1(X6, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X6, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X6)) => ((g3_lattices(X1, sK5(X0, X1, X2), sK6(X0, X1, X2)) = X0) & (k1_realset1(u1_lattices(X2), X1) = sK6(X0, X1, X2)) & (k1_realset1(u2_lattices(X2), X1) = sK5(X0, X1, X2)) & m2_relset_1(sK6(X0, X1, X2), k2_zfmisc_1(X1, X1), X1) & v1_funct_2(sK6(X0, X1, X2), k2_zfmisc_1(X1, X1), X1) & v1_funct_1(sK6(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f237, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | ! [X3] : (! [X4] : (~ (g3_lattices(X1, X3, X4) = X0) | ~ (k1_realset1(u1_lattices(X2), X1) = X4) | ~ (k1_realset1(u2_lattices(X2), X1) = X3) | ~ m2_relset_1(X4, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_2(X4, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_1(X4)) | ~ m2_relset_1(X3, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_2(X3, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_1(X3))) & (? [X5] : (? [X6] : ((g3_lattices(X1, X5, X6) = X0) & (k1_realset1(u1_lattices(X2), X1) = X6) & (k1_realset1(u2_lattices(X2), X1) = X5) & m2_relset_1(X6, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X6, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X6)) & m2_relset_1(X5, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X5, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X5)) | ~ sP0(X0, X1, X2))), inference(rectify, [], [f236])).
fof(f236, plain, ! [X2, X1, X0] : ((sP0(X2, X1, X0) | ! [X3] : (! [X4] : (~ (g3_lattices(X1, X3, X4) = X2) | ~ (k1_realset1(u1_lattices(X0), X1) = X4) | ~ (k1_realset1(u2_lattices(X0), X1) = X3) | ~ m2_relset_1(X4, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_2(X4, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_1(X4)) | ~ m2_relset_1(X3, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_2(X3, k2_zfmisc_1(X1, X1), X1) | ~ v1_funct_1(X3))) & (? [X3] : (? [X4] : ((g3_lattices(X1, X3, X4) = X2) & (k1_realset1(u1_lattices(X0), X1) = X4) & (k1_realset1(u2_lattices(X0), X1) = X3) & m2_relset_1(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X4)) & m2_relset_1(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X3)) | ~ sP0(X2, X1, X0))), inference(nnf_transformation, [], [f226])).
fof(f226, plain, ! [X2, X1, X0] : (sP0(X2, X1, X0) <=> ? [X3] : (? [X4] : ((g3_lattices(X1, X3, X4) = X2) & (k1_realset1(u1_lattices(X0), X1) = X4) & (k1_realset1(u2_lattices(X0), X1) = X3) & m2_relset_1(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X4)) & m2_relset_1(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X3))), inference(usedef, [], [e226])).
fof(e226, plain, ! [X2, X1, X0] : (sP0(X2, X1, X0) <=> ? [X3] : (? [X4] : ((g3_lattices(X1, X3, X4) = X2) & (k1_realset1(u1_lattices(X0), X1) = X4) & (k1_realset1(u2_lattices(X0), X1) = X3) & m2_relset_1(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X4)) & m2_relset_1(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f1103, plain, ! [X0, X1] : (sP0(k23_filter_2(X0, X1), X1, X0) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_lattice4(X1, X0)), inference(subsumption_resolution, [], [f1101, f347])).
fof(f347, plain, ! [X0, X1] : (m2_nat_lat(k23_filter_2(X0, X1), X0) | ~ m2_lattice4(X1, X0) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ! [X0, X1] : (m2_nat_lat(k23_filter_2(X0, X1), X0) | ~ m2_lattice4(X1, X0) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f142])).
fof(f142, plain, ! [X0, X1] : (m2_nat_lat(k23_filter_2(X0, X1), X0) | (~ m2_lattice4(X1, X0) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ! [X0, X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m2_nat_lat(k23_filter_2(X0, X1), X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', dt_k23_filter_2)).
fof(f1101, plain, ! [X0, X1] : (~ m2_nat_lat(k23_filter_2(X0, X1), X0) | ~ m2_lattice4(X1, X0) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | sP0(k23_filter_2(X0, X1), X1, X0)), inference(resolution, [], [f332, f487])).
fof(f487, plain, ! [X0, X1] : (~ sP1(X0, X1, k23_filter_2(X0, X1)) | sP0(k23_filter_2(X0, X1), X1, X0)), inference(equality_resolution, [], [f320])).
fof(f320, plain, ! [X2, X0, X1] : (sP0(X2, X1, X0) | ~ (k23_filter_2(X0, X1) = X2) | ~ sP1(X0, X1, X2)), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ! [X0, X1, X2] : ((((k23_filter_2(X0, X1) = X2) | ~ sP0(X2, X1, X0)) & (sP0(X2, X1, X0) | ~ (k23_filter_2(X0, X1) = X2))) | ~ sP1(X0, X1, X2)), inference(nnf_transformation, [], [f227])).
fof(f227, plain, ! [X0, X1, X2] : (((k23_filter_2(X0, X1) = X2) <=> sP0(X2, X1, X0)) | ~ sP1(X0, X1, X2)), inference(usedef, [], [e227])).
fof(e227, plain, ! [X0, X1, X2] : (sP1(X0, X1, X2) <=> ((k23_filter_2(X0, X1) = X2) <=> sP0(X2, X1, X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f332, plain, ! [X2, X0, X1] : (sP1(X0, X1, X2) | ~ m2_nat_lat(X2, X0) | ~ m2_lattice4(X1, X0) | v1_xboole_0(X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f228])).
fof(f228, plain, ! [X0] : (! [X1] : (! [X2] : (sP1(X0, X1, X2) | ~ m2_nat_lat(X2, X0)) | ~ m2_lattice4(X1, X0) | v1_xboole_0(X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(definition_folding, [], [f126, e227, e226])).
fof(f126, plain, ! [X0] : (! [X1] : (! [X2] : (((k23_filter_2(X0, X1) = X2) <=> ? [X3] : (? [X4] : ((g3_lattices(X1, X3, X4) = X2) & (k1_realset1(u1_lattices(X0), X1) = X4) & (k1_realset1(u2_lattices(X0), X1) = X3) & m2_relset_1(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X4)) & m2_relset_1(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X3))) | ~ m2_nat_lat(X2, X0)) | ~ m2_lattice4(X1, X0) | v1_xboole_0(X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f125])).
fof(f125, plain, ! [X0] : (! [X1] : (! [X2] : (((k23_filter_2(X0, X1) = X2) <=> ? [X3] : (? [X4] : ((g3_lattices(X1, X3, X4) = X2) & (k1_realset1(u1_lattices(X0), X1) = X4) & (k1_realset1(u2_lattices(X0), X1) = X3) & m2_relset_1(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X4)) & m2_relset_1(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X3))) | ~ m2_nat_lat(X2, X0)) | (~ m2_lattice4(X1, X0) | v1_xboole_0(X1))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) => ! [X2] : (m2_nat_lat(X2, X0) => ((k23_filter_2(X0, X1) = X2) <=> ? [X3] : (? [X4] : ((g3_lattices(X1, X3, X4) = X2) & (k1_realset1(u1_lattices(X0), X1) = X4) & (k1_realset1(u2_lattices(X0), X1) = X3) & m2_relset_1(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X4, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X4)) & m2_relset_1(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_2(X3, k2_zfmisc_1(X1, X1), X1) & v1_funct_1(X3)))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', d16_filter_2)).
fof(f45033, plain, ((k23_filter_2(sK4, k17_filter_2(sK4)) = g3_lattices(k17_filter_2(sK4), sK5(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4), u1_lattices(sK4))) | (~ spl33_72 | ~ spl33_100)), inference(forward_demodulation, [], [f45032, f24239])).
fof(f24239, plain, ((u1_lattices(sK4) = sK6(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4)) | (~ spl33_72 | ~ spl33_100)), inference(forward_demodulation, [], [f24238, f4104])).
fof(f4104, plain, ((u1_lattices(sK4) = k1_realset1(u1_lattices(sK4), k17_filter_2(sK4))) | (~ spl33_72 | ~ spl33_100)), inference(forward_demodulation, [], [f4099, f3282])).
fof(f3282, plain, ((u1_lattices(sK4) = k7_relat_1(u1_lattices(sK4), k1_relat_1(u2_lattices(sK4)))) | (~ spl33_72 | ~ spl33_100)), inference(backward_demodulation, [], [f3188, f3281])).
fof(f3188, plain, (u1_lattices(sK4) = k7_relat_1(u1_lattices(sK4), k1_relat_1(u1_lattices(sK4)))), inference(resolution, [], [f3176, f828])).
fof(f3176, plain, v1_relat_1(u1_lattices(sK4)), inference(resolution, [], [f875, f684])).
fof(f875, plain, m2_relset_1(u1_lattices(sK4), k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4)), inference(subsumption_resolution, [], [f873, f513])).
fof(f873, plain, (m2_relset_1(u1_lattices(sK4), k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4)) | ~ l1_lattices(sK4)), inference(superposition, [], [f368, f755])).
fof(f4099, plain, ((k7_relat_1(u1_lattices(sK4), k1_relat_1(u2_lattices(sK4))) = k1_realset1(u1_lattices(sK4), k17_filter_2(sK4))) | (~ spl33_72 | ~ spl33_100)), inference(superposition, [], [f2086, f3283])).
fof(f2086, plain, ! [X0] : (k1_realset1(u1_lattices(sK4), X0) = k7_relat_1(u1_lattices(sK4), k2_zfmisc_1(X0, X0))), inference(subsumption_resolution, [], [f2085, f301])).
fof(f2085, plain, ! [X0] : ((k1_realset1(u1_lattices(sK4), X0) = k7_relat_1(u1_lattices(sK4), k2_zfmisc_1(X0, X0))) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f2079, f613])).
fof(f613, plain, v7_lattices(sK4), inference(subsumption_resolution, [], [f612, f301])).
fof(f612, plain, (v3_struct_0(sK4) | v7_lattices(sK4)), inference(subsumption_resolution, [], [f607, f302])).
fof(f607, plain, (~ v10_lattices(sK4) | v3_struct_0(sK4) | v7_lattices(sK4)), inference(resolution, [], [f312, f303])).
fof(f312, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v7_lattices(X0)), inference(cnf_transformation, [], [f119])).
fof(f2079, plain, ! [X0] : ((k1_realset1(u1_lattices(sK4), X0) = k7_relat_1(u1_lattices(sK4), k2_zfmisc_1(X0, X0))) | ~ v7_lattices(sK4) | v3_struct_0(sK4)), inference(resolution, [], [f738, f513])).
fof(f738, plain, ! [X4, X5] : (~ l1_lattices(X4) | (k1_realset1(u1_lattices(X4), X5) = k7_relat_1(u1_lattices(X4), k2_zfmisc_1(X5, X5))) | ~ v7_lattices(X4) | v3_struct_0(X4)), inference(resolution, [], [f340, f414])).
fof(f414, plain, ! [X0] : (v1_relat_1(u1_lattices(X0)) | ~ l1_lattices(X0) | ~ v7_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f193])).
fof(f193, plain, ! [X0] : ((v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)) & v1_relat_1(u1_lattices(X0))) | ~ l1_lattices(X0) | ~ v7_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f192])).
fof(f192, plain, ! [X0] : ((v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)) & v1_relat_1(u1_lattices(X0))) | (~ l1_lattices(X0) | ~ v7_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f105])).
fof(f105, plain, ! [X0] : ((l1_lattices(X0) & v7_lattices(X0) & ~ v3_struct_0(X0)) => (v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)) & v1_relat_1(u1_lattices(X0)))), inference(pure_predicate_removal, [], [f103])).
fof(f103, plain, ! [X0] : ((l1_lattices(X0) & v7_lattices(X0) & ~ v3_struct_0(X0)) => (v1_partfun1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)) & v1_relat_1(u1_lattices(X0)))), inference(pure_predicate_removal, [], [f67])).
fof(f67, plain, ! [X0] : ((l1_lattices(X0) & v7_lattices(X0) & ~ v3_struct_0(X0)) => (v1_partfun1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v2_binop_1(u1_lattices(X0), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)) & v1_relat_1(u1_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', fc5_lattice2)).
fof(f24238, plain, (k1_realset1(u1_lattices(sK4), k17_filter_2(sK4)) = sK6(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4)), inference(subsumption_resolution, [], [f24237, f301])).
fof(f24237, plain, (v3_struct_0(sK4) | (k1_realset1(u1_lattices(sK4), k17_filter_2(sK4)) = sK6(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4))), inference(subsumption_resolution, [], [f24224, f302])).
fof(f24224, plain, (~ v10_lattices(sK4) | v3_struct_0(sK4) | (k1_realset1(u1_lattices(sK4), k17_filter_2(sK4)) = sK6(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4))), inference(resolution, [], [f6305, f303])).
fof(f6305, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | (k1_realset1(u1_lattices(X0), k17_filter_2(X0)) = sK6(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0))), inference(subsumption_resolution, [], [f6304, f788])).
fof(f6304, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v1_xboole_0(k17_filter_2(X0)) | (k1_realset1(u1_lattices(X0), k17_filter_2(X0)) = sK6(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0))), inference(duplicate_literal_removal, [], [f6286])).
fof(f6286, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v1_xboole_0(k17_filter_2(X0)) | (k1_realset1(u1_lattices(X0), k17_filter_2(X0)) = sK6(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f2166, f859])).
fof(f2166, plain, ! [X2, X3] : (~ m2_lattice4(X2, X3) | ~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | v1_xboole_0(X2) | (k1_realset1(u1_lattices(X3), X2) = sK6(k23_filter_2(X3, X2), X2, X3))), inference(resolution, [], [f1103, f329])).
fof(f329, plain, ! [X2, X0, X1] : (~ sP0(X0, X1, X2) | (k1_realset1(u1_lattices(X2), X1) = sK6(X0, X1, X2))), inference(cnf_transformation, [], [f240])).
fof(f45032, plain, (k23_filter_2(sK4, k17_filter_2(sK4)) = g3_lattices(k17_filter_2(sK4), sK5(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4), sK6(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4))), inference(subsumption_resolution, [], [f45031, f301])).
fof(f45031, plain, (v3_struct_0(sK4) | (k23_filter_2(sK4, k17_filter_2(sK4)) = g3_lattices(k17_filter_2(sK4), sK5(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4), sK6(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4)))), inference(subsumption_resolution, [], [f45018, f302])).
fof(f45018, plain, (~ v10_lattices(sK4) | v3_struct_0(sK4) | (k23_filter_2(sK4, k17_filter_2(sK4)) = g3_lattices(k17_filter_2(sK4), sK5(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4), sK6(k23_filter_2(sK4, k17_filter_2(sK4)), k17_filter_2(sK4), sK4)))), inference(resolution, [], [f7018, f303])).
fof(f7018, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | (k23_filter_2(X0, k17_filter_2(X0)) = g3_lattices(k17_filter_2(X0), sK5(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0), sK6(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0)))), inference(subsumption_resolution, [], [f7017, f788])).
fof(f7017, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v1_xboole_0(k17_filter_2(X0)) | (k23_filter_2(X0, k17_filter_2(X0)) = g3_lattices(k17_filter_2(X0), sK5(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0), sK6(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0)))), inference(duplicate_literal_removal, [], [f6999])).
fof(f6999, plain, ! [X0] : (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | v1_xboole_0(k17_filter_2(X0)) | (k23_filter_2(X0, k17_filter_2(X0)) = g3_lattices(k17_filter_2(X0), sK5(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0), sK6(k23_filter_2(X0, k17_filter_2(X0)), k17_filter_2(X0), X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f2165, f859])).
fof(f2165, plain, ! [X0, X1] : (~ m2_lattice4(X0, X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | v1_xboole_0(X0) | (k23_filter_2(X1, X0) = g3_lattices(X0, sK5(k23_filter_2(X1, X0), X0, X1), sK6(k23_filter_2(X1, X0), X0, X1)))), inference(resolution, [], [f1103, f330])).
fof(f330, plain, ! [X2, X0, X1] : (~ sP0(X0, X1, X2) | (g3_lattices(X1, sK5(X0, X1, X2), sK6(X0, X1, X2)) = X0)), inference(cnf_transformation, [], [f240])).
fof(f45112, plain, (~ (k23_filter_2(sK4, k17_filter_2(sK4)) = g3_lattices(k17_filter_2(sK4), u2_lattices(sK4), u1_lattices(sK4))) | spl33_2), inference(forward_demodulation, [], [f45111, f755])).
fof(f45111, plain, (~ (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k17_filter_2(sK4))) | spl33_2), inference(forward_demodulation, [], [f505, f817])).
fof(f817, plain, (k1_filter_2(sK4) = k17_filter_2(sK4)), inference(forward_demodulation, [], [f816, f757])).
fof(f757, plain, (k17_filter_2(sK4) = k1_filter_0(sK4)), inference(backward_demodulation, [], [f729, f755])).
fof(f729, plain, (u1_struct_0(sK4) = k1_filter_0(sK4)), inference(subsumption_resolution, [], [f728, f301])).
fof(f728, plain, ((u1_struct_0(sK4) = k1_filter_0(sK4)) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f723, f302])).
fof(f723, plain, ((u1_struct_0(sK4) = k1_filter_0(sK4)) | ~ v10_lattices(sK4) | v3_struct_0(sK4)), inference(resolution, [], [f339, f303])).
fof(f339, plain, ! [X0] : (~ l3_lattices(X0) | (u1_struct_0(X0) = k1_filter_0(X0)) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ! [X0] : ((u1_struct_0(X0) = k1_filter_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f129])).
fof(f129, plain, ! [X0] : ((u1_struct_0(X0) = k1_filter_0(X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (u1_struct_0(X0) = k1_filter_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', d2_filter_0)).
fof(f816, plain, (k1_filter_2(sK4) = k1_filter_0(sK4)), inference(subsumption_resolution, [], [f815, f301])).
fof(f815, plain, ((k1_filter_2(sK4) = k1_filter_0(sK4)) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f810, f302])).
fof(f810, plain, ((k1_filter_2(sK4) = k1_filter_0(sK4)) | ~ v10_lattices(sK4) | v3_struct_0(sK4)), inference(resolution, [], [f465, f303])).
fof(f465, plain, ! [X0] : (~ l3_lattices(X0) | (k1_filter_2(X0) = k1_filter_0(X0)) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ! [X0] : ((k1_filter_2(X0) = k1_filter_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f203])).
fof(f203, plain, ! [X0] : ((k1_filter_2(X0) = k1_filter_0(X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k1_filter_2(X0) = k1_filter_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', redefinition_k1_filter_2)).
fof(f505, plain, (~ (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k1_filter_2(sK4))) | spl33_2), inference(avatar_component_clause, [], [f503])).
fof(f503, plain, (spl33_2 <=> (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k1_filter_2(sK4)))), introduced(avatar_definition, [new_symbols(naming, [spl33_2])])).
fof(f45036, plain, (spl33_1 | ~ spl33_72 | ~ spl33_100), inference(avatar_contradiction_clause, [], [f45035])).
fof(f45035, plain, ($false | (spl33_1 | ~ spl33_72 | ~ spl33_100)), inference(subsumption_resolution, [], [f45034, f758])).
fof(f758, plain, (~ (k23_filter_2(sK4, k17_filter_2(sK4)) = g3_lattices(k17_filter_2(sK4), u2_lattices(sK4), u1_lattices(sK4))) | spl33_1), inference(backward_demodulation, [], [f501, f755])).
fof(f501, plain, (~ (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k17_filter_2(sK4))) | spl33_1), inference(avatar_component_clause, [], [f499])).
fof(f499, plain, (spl33_1 <=> (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k17_filter_2(sK4)))), introduced(avatar_definition, [new_symbols(naming, [spl33_1])])).
fof(f2487, plain, ~ spl33_73, inference(avatar_contradiction_clause, [], [f2486])).
fof(f2486, plain, ($false | ~ spl33_73), inference(subsumption_resolution, [], [f2445, f387])).
fof(f387, plain, v1_xboole_0(k1_xboole_0), inference(cnf_transformation, [], [f58])).
fof(f58, plain, v1_xboole_0(k1_xboole_0), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', fc1_xboole_0)).
fof(f2445, plain, (~ v1_xboole_0(k1_xboole_0) | ~ spl33_73), inference(backward_demodulation, [], [f774, f2197])).
fof(f2197, plain, ((k1_xboole_0 = k17_filter_2(sK4)) | ~ spl33_73), inference(avatar_component_clause, [], [f2195])).
fof(f2195, plain, (spl33_73 <=> (k1_xboole_0 = k17_filter_2(sK4))), introduced(avatar_definition, [new_symbols(naming, [spl33_73])])).
fof(f774, plain, ~ v1_xboole_0(k17_filter_2(sK4)), inference(subsumption_resolution, [], [f773, f301])).
fof(f773, plain, (~ v1_xboole_0(k17_filter_2(sK4)) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f770, f518])).
fof(f518, plain, l1_struct_0(sK4), inference(resolution, [], [f513, f351])).
fof(f351, plain, ! [X0] : (~ l1_lattices(X0) | l1_struct_0(X0)), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ! [X0] : (l1_struct_0(X0) | ~ l1_lattices(X0)), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ! [X0] : (l1_lattices(X0) => l1_struct_0(X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', dt_l1_lattices)).
fof(f770, plain, (~ v1_xboole_0(k17_filter_2(sK4)) | ~ l1_struct_0(sK4) | v3_struct_0(sK4)), inference(superposition, [], [f385, f755])).
fof(f385, plain, ! [X0] : (~ v1_xboole_0(u1_struct_0(X0)) | ~ l1_struct_0(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ! [X0] : (~ v1_xboole_0(u1_struct_0(X0)) | ~ l1_struct_0(X0) | v3_struct_0(X0)), inference(flattening, [], [f174])).
fof(f174, plain, ! [X0] : (~ v1_xboole_0(u1_struct_0(X0)) | (~ l1_struct_0(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f56])).
fof(f56, plain, ! [X0] : ((l1_struct_0(X0) & ~ v3_struct_0(X0)) => ~ v1_xboole_0(u1_struct_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', fc1_struct_0)).
fof(f2376, plain, (spl33_100 | spl33_73), inference(avatar_split_clause, [], [f2371, f2195, f2373])).
fof(f2371, plain, ((k1_xboole_0 = k17_filter_2(sK4)) | (k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)) = k4_relset_1(k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4), u2_lattices(sK4)))), inference(forward_demodulation, [], [f2370, f755])).
fof(f2370, plain, ((k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)) = k4_relset_1(k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4), u2_lattices(sK4))) | (k1_xboole_0 = u1_struct_0(sK4))), inference(forward_demodulation, [], [f2364, f755])).
fof(f2364, plain, ((k2_zfmisc_1(u1_struct_0(sK4), u1_struct_0(sK4)) = k4_relset_1(k2_zfmisc_1(u1_struct_0(sK4), u1_struct_0(sK4)), u1_struct_0(sK4), u2_lattices(sK4))) | (k1_xboole_0 = u1_struct_0(sK4))), inference(resolution, [], [f1068, f519])).
fof(f1068, plain, ! [X9] : (~ l2_lattices(X9) | (k2_zfmisc_1(u1_struct_0(X9), u1_struct_0(X9)) = k4_relset_1(k2_zfmisc_1(u1_struct_0(X9), u1_struct_0(X9)), u1_struct_0(X9), u2_lattices(X9))) | (k1_xboole_0 = u1_struct_0(X9))), inference(subsumption_resolution, [], [f1062, f370])).
fof(f370, plain, ! [X0] : (v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f162])).
fof(f1062, plain, ! [X9] : (~ v1_funct_2(u2_lattices(X9), k2_zfmisc_1(u1_struct_0(X9), u1_struct_0(X9)), u1_struct_0(X9)) | (k1_xboole_0 = u1_struct_0(X9)) | (k2_zfmisc_1(u1_struct_0(X9), u1_struct_0(X9)) = k4_relset_1(k2_zfmisc_1(u1_struct_0(X9), u1_struct_0(X9)), u1_struct_0(X9), u2_lattices(X9))) | ~ l2_lattices(X9)), inference(resolution, [], [f333, f371])).
fof(f333, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | ~ v1_funct_2(X2, X0, X1) | (k1_xboole_0 = X1) | (k4_relset_1(X0, X1, X2) = X0)), inference(cnf_transformation, [], [f241])).
fof(f241, plain, ! [X0, X1, X2] : (((((v1_funct_2(X2, X0, X1) | ~ (k1_xboole_0 = X2)) & ((k1_xboole_0 = X2) | ~ v1_funct_2(X2, X0, X1))) | (k1_xboole_0 = X0) | ~ (k1_xboole_0 = X1)) & (((v1_funct_2(X2, X0, X1) | ~ (k4_relset_1(X0, X1, X2) = X0)) & ((k4_relset_1(X0, X1, X2) = X0) | ~ v1_funct_2(X2, X0, X1))) | (~ (k1_xboole_0 = X0) & (k1_xboole_0 = X1)))) | ~ m2_relset_1(X2, X0, X1)), inference(nnf_transformation, [], [f128])).
fof(f128, plain, ! [X0, X1, X2] : ((((v1_funct_2(X2, X0, X1) <=> (k1_xboole_0 = X2)) | (k1_xboole_0 = X0) | ~ (k1_xboole_0 = X1)) & ((v1_funct_2(X2, X0, X1) <=> (k4_relset_1(X0, X1, X2) = X0)) | (~ (k1_xboole_0 = X0) & (k1_xboole_0 = X1)))) | ~ m2_relset_1(X2, X0, X1)), inference(flattening, [], [f127])).
fof(f127, plain, ! [X0, X1, X2] : (((((v1_funct_2(X2, X0, X1) <=> (k1_xboole_0 = X2)) | (k1_xboole_0 = X0)) | ~ (k1_xboole_0 = X1)) & ((v1_funct_2(X2, X0, X1) <=> (k4_relset_1(X0, X1, X2) = X0)) | (~ (k1_xboole_0 = X0) & (k1_xboole_0 = X1)))) | ~ m2_relset_1(X2, X0, X1)), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) => (((k1_xboole_0 = X1) => ((v1_funct_2(X2, X0, X1) <=> (k1_xboole_0 = X2)) | (k1_xboole_0 = X0))) & (((k1_xboole_0 = X1) => (k1_xboole_0 = X0)) => (v1_funct_2(X2, X0, X1) <=> (k4_relset_1(X0, X1, X2) = X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT332+1.p', d1_funct_2)).
fof(f2198, plain, (spl33_72 | spl33_73), inference(avatar_split_clause, [], [f2189, f2195, f2191])).
fof(f2189, plain, ((k1_xboole_0 = k17_filter_2(sK4)) | (k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)) = k4_relset_1(k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4), u1_lattices(sK4)))), inference(forward_demodulation, [], [f2188, f755])).
fof(f2188, plain, ((k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)) = k4_relset_1(k2_zfmisc_1(k17_filter_2(sK4), k17_filter_2(sK4)), k17_filter_2(sK4), u1_lattices(sK4))) | (k1_xboole_0 = u1_struct_0(sK4))), inference(forward_demodulation, [], [f2182, f755])).
fof(f2182, plain, ((k2_zfmisc_1(u1_struct_0(sK4), u1_struct_0(sK4)) = k4_relset_1(k2_zfmisc_1(u1_struct_0(sK4), u1_struct_0(sK4)), u1_struct_0(sK4), u1_lattices(sK4))) | (k1_xboole_0 = u1_struct_0(sK4))), inference(resolution, [], [f1067, f513])).
fof(f1067, plain, ! [X8] : (~ l1_lattices(X8) | (k2_zfmisc_1(u1_struct_0(X8), u1_struct_0(X8)) = k4_relset_1(k2_zfmisc_1(u1_struct_0(X8), u1_struct_0(X8)), u1_struct_0(X8), u1_lattices(X8))) | (k1_xboole_0 = u1_struct_0(X8))), inference(subsumption_resolution, [], [f1061, f367])).
fof(f367, plain, ! [X0] : (v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f161])).
fof(f1061, plain, ! [X8] : (~ v1_funct_2(u1_lattices(X8), k2_zfmisc_1(u1_struct_0(X8), u1_struct_0(X8)), u1_struct_0(X8)) | (k1_xboole_0 = u1_struct_0(X8)) | (k2_zfmisc_1(u1_struct_0(X8), u1_struct_0(X8)) = k4_relset_1(k2_zfmisc_1(u1_struct_0(X8), u1_struct_0(X8)), u1_struct_0(X8), u1_lattices(X8))) | ~ l1_lattices(X8)), inference(resolution, [], [f333, f368])).
fof(f506, plain, (~ spl33_1 | ~ spl33_2), inference(avatar_split_clause, [], [f304, f503, f499])).
fof(f304, plain, (~ (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k1_filter_2(sK4))) | ~ (g3_lattices(u1_struct_0(sK4), u2_lattices(sK4), u1_lattices(sK4)) = k23_filter_2(sK4, k17_filter_2(sK4)))), inference(cnf_transformation, [], [f234])).