fof(f10700, plain, $false, inference(avatar_sat_refutation, [], [f547, f548, f1136, f1147, f1159, f1231, f10640, f10693])).
fof(f10693, plain, (~ spl33_1 | ~ spl33_2 | spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35), inference(avatar_contradiction_clause, [], [f10692])).
fof(f10692, plain, ($false | (~ spl33_1 | ~ spl33_2 | spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f10691, f2110])).
fof(f2110, plain, (m1_filter_0(sK3, k1_lattice2(sK2)) | (spl33_29 | ~ spl33_30 | ~ spl33_31)), inference(backward_demodulation, [], [f1826, f2106])).
fof(f2106, plain, (sK3 = k7_filter_2(sK2, sK3)), inference(subsumption_resolution, [], [f2105, f295])).
fof(f295, plain, ~ v3_struct_0(sK2), inference(cnf_transformation, [], [f232])).
fof(f232, plain, ((((r2_hidden(k7_lattices(sK2, sK4), sK3) | ~ r2_hidden(sK4, sK3)) & (~ r2_hidden(k7_lattices(sK2, sK4), sK3) | r2_hidden(sK4, sK3)) & m1_subset_1(sK4, u1_struct_0(sK2))) & r2_filter_2(sK2, sK3) & m2_filter_2(sK3, sK2)) & l3_lattices(sK2) & v17_lattices(sK2) & v10_lattices(sK2) & ~ v3_struct_0(sK2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2, sK3, sK4])], [f228, f231, f230, f229])).
fof(f229, plain, (? [X0] : (? [X1] : (? [X2] : ((r2_hidden(k7_lattices(X0, X2), X1) | ~ r2_hidden(X2, X1)) & (~ r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1)) & m1_subset_1(X2, u1_struct_0(X0))) & r2_filter_2(X0, X1) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (? [X1] : (? [X2] : ((r2_hidden(k7_lattices(sK2, X2), X1) | ~ r2_hidden(X2, X1)) & (~ r2_hidden(k7_lattices(sK2, X2), X1) | r2_hidden(X2, X1)) & m1_subset_1(X2, u1_struct_0(sK2))) & r2_filter_2(sK2, X1) & m2_filter_2(X1, sK2)) & l3_lattices(sK2) & v17_lattices(sK2) & v10_lattices(sK2) & ~ v3_struct_0(sK2))), introduced(choice_axiom, [])).
fof(f230, plain, (? [X1] : (? [X2] : ((r2_hidden(k7_lattices(sK2, X2), X1) | ~ r2_hidden(X2, X1)) & (~ r2_hidden(k7_lattices(sK2, X2), X1) | r2_hidden(X2, X1)) & m1_subset_1(X2, u1_struct_0(sK2))) & r2_filter_2(sK2, X1) & m2_filter_2(X1, sK2)) => (? [X2] : ((r2_hidden(k7_lattices(sK2, X2), sK3) | ~ r2_hidden(X2, sK3)) & (~ r2_hidden(k7_lattices(sK2, X2), sK3) | r2_hidden(X2, sK3)) & m1_subset_1(X2, u1_struct_0(sK2))) & r2_filter_2(sK2, sK3) & m2_filter_2(sK3, sK2))), introduced(choice_axiom, [])).
fof(f231, plain, (? [X2] : ((r2_hidden(k7_lattices(sK2, X2), sK3) | ~ r2_hidden(X2, sK3)) & (~ r2_hidden(k7_lattices(sK2, X2), sK3) | r2_hidden(X2, sK3)) & m1_subset_1(X2, u1_struct_0(sK2))) => ((r2_hidden(k7_lattices(sK2, sK4), sK3) | ~ r2_hidden(sK4, sK3)) & (~ r2_hidden(k7_lattices(sK2, sK4), sK3) | r2_hidden(sK4, sK3)) & m1_subset_1(sK4, u1_struct_0(sK2)))), introduced(choice_axiom, [])).
fof(f228, plain, ? [X0] : (? [X1] : (? [X2] : ((r2_hidden(k7_lattices(X0, X2), X1) | ~ r2_hidden(X2, X1)) & (~ r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1)) & m1_subset_1(X2, u1_struct_0(X0))) & r2_filter_2(X0, X1) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f227])).
fof(f227, plain, ? [X0] : (? [X1] : (? [X2] : (((r2_hidden(k7_lattices(X0, X2), X1) | ~ r2_hidden(X2, X1)) & (~ r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1))) & m1_subset_1(X2, u1_struct_0(X0))) & r2_filter_2(X0, X1) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(nnf_transformation, [], [f114])).
fof(f114, plain, ? [X0] : (? [X1] : (? [X2] : (~ (r2_hidden(X2, X1) <=> ~ r2_hidden(k7_lattices(X0, X2), X1)) & m1_subset_1(X2, u1_struct_0(X0))) & r2_filter_2(X0, X1) & m2_filter_2(X1, X0)) & l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f113])).
fof(f113, plain, ? [X0] : (? [X1] : ((? [X2] : (~ (r2_hidden(X2, X1) <=> ~ r2_hidden(k7_lattices(X0, X2), X1)) & m1_subset_1(X2, u1_struct_0(X0))) & r2_filter_2(X0, X1)) & m2_filter_2(X1, X0)) & (l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (r2_filter_2(X0, X1) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (r2_hidden(X2, X1) <=> ~ r2_hidden(k7_lattices(X0, X2), X1)))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (r2_filter_2(X0, X1) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (r2_hidden(X2, X1) <=> ~ r2_hidden(k7_lattices(X0, X2), X1)))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', t59_filter_2)).
fof(f2105, plain, (v3_struct_0(sK2) | (sK3 = k7_filter_2(sK2, sK3))), inference(subsumption_resolution, [], [f2104, f296])).
fof(f296, plain, v10_lattices(sK2), inference(cnf_transformation, [], [f232])).
fof(f2104, plain, (~ v10_lattices(sK2) | v3_struct_0(sK2) | (sK3 = k7_filter_2(sK2, sK3))), inference(subsumption_resolution, [], [f2098, f298])).
fof(f298, plain, l3_lattices(sK2), inference(cnf_transformation, [], [f232])).
fof(f2098, plain, (~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2) | (sK3 = k7_filter_2(sK2, sK3))), inference(resolution, [], [f1071, f1000])).
fof(f1000, plain, m2_lattice4(sK3, sK2), inference(subsumption_resolution, [], [f999, f295])).
fof(f999, plain, (m2_lattice4(sK3, sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f998, f296])).
fof(f998, plain, (m2_lattice4(sK3, sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f995, f298])).
fof(f995, plain, (m2_lattice4(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f360, f299])).
fof(f299, plain, m2_filter_2(sK3, sK2), inference(cnf_transformation, [], [f232])).
fof(f360, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f159])).
fof(f159, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', dt_m2_filter_2)).
fof(f1071, plain, ! [X0, X1] : (~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | (k7_filter_2(X0, X1) = X1)), inference(duplicate_literal_removal, [], [f1062])).
fof(f1062, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f342, f361])).
fof(f361, plain, ! [X0, X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f161])).
fof(f161, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_lattice4(X1, X0) => m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', dt_m2_lattice4)).
fof(f342, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | (k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f139])).
fof(f139, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) => (k7_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', d6_filter_2)).
fof(f1826, plain, (m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | (spl33_29 | ~ spl33_30 | ~ spl33_31)), inference(subsumption_resolution, [], [f1825, f1094])).
fof(f1094, plain, (~ v3_struct_0(k1_lattice2(sK2)) | spl33_29), inference(avatar_component_clause, [], [f1093])).
fof(f1093, plain, (spl33_29 <=> v3_struct_0(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl33_29])])).
fof(f1825, plain, (m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | (~ spl33_30 | ~ spl33_31)), inference(subsumption_resolution, [], [f1824, f1098])).
fof(f1098, plain, (v10_lattices(k1_lattice2(sK2)) | ~ spl33_30), inference(avatar_component_clause, [], [f1097])).
fof(f1097, plain, (spl33_30 <=> v10_lattices(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl33_30])])).
fof(f1824, plain, (m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2)) | ~ spl33_31), inference(subsumption_resolution, [], [f1821, f1102])).
fof(f1102, plain, (l3_lattices(k1_lattice2(sK2)) | ~ spl33_31), inference(avatar_component_clause, [], [f1101])).
fof(f1101, plain, (spl33_31 <=> l3_lattices(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl33_31])])).
fof(f1821, plain, (m1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))), inference(resolution, [], [f1317, f522])).
fof(f522, plain, ! [X0, X1] : (~ m1_filter_2(X1, X0) | m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f291])).
fof(f291, plain, ! [X0] : (! [X1] : ((m1_filter_2(X1, X0) | ~ m1_filter_0(X1, X0)) & (m1_filter_0(X1, X0) | ~ m1_filter_2(X1, X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f208])).
fof(f208, plain, ! [X0] : (! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f207])).
fof(f207, plain, ! [X0] : (! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f85])).
fof(f85, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', redefinition_m1_filter_2)).
fof(f1317, plain, m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)), inference(subsumption_resolution, [], [f1316, f295])).
fof(f1316, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1315, f296])).
fof(f1315, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1314, f298])).
fof(f1314, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1308, f299])).
fof(f1308, plain, (m1_filter_2(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f345, f1080])).
fof(f1080, plain, (k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)), inference(subsumption_resolution, [], [f1079, f295])).
fof(f1079, plain, ((k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1078, f296])).
fof(f1078, plain, ((k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1075, f298])).
fof(f1075, plain, ((k7_filter_2(sK2, sK3) = k15_filter_2(sK2, sK3)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f521, f299])).
fof(f521, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | (k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f205])).
fof(f205, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | (~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f84])).
fof(f84, plain, ! [X0, X1] : ((m2_filter_2(X1, X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k7_filter_2(X0, X1) = k15_filter_2(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', redefinition_k15_filter_2)).
fof(f345, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f143])).
fof(f143, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | (~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ! [X0, X1] : ((m2_filter_2(X1, X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', dt_k15_filter_2)).
fof(f10691, plain, (~ m1_filter_0(sK3, k1_lattice2(sK2)) | (~ spl33_1 | ~ spl33_2 | spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f10690, f2108])).
fof(f2108, plain, v1_filter_0(sK3, k1_lattice2(sK2)), inference(backward_demodulation, [], [f1313, f2106])).
fof(f1313, plain, v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)), inference(subsumption_resolution, [], [f1312, f295])).
fof(f1312, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1311, f296])).
fof(f1311, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1310, f298])).
fof(f1310, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1309, f299])).
fof(f1309, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1307, f300])).
fof(f300, plain, r2_filter_2(sK2, sK3), inference(cnf_transformation, [], [f232])).
fof(f1307, plain, (v1_filter_0(k7_filter_2(sK2, sK3), k1_lattice2(sK2)) | ~ r2_filter_2(sK2, sK3) | ~ m2_filter_2(sK3, sK2) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f529, f1080])).
fof(f529, plain, ! [X0, X1] : (v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ r2_filter_2(X0, X1) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f293])).
fof(f293, plain, ! [X0] : (! [X1] : (((r2_filter_2(X0, X1) | ~ v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) & (v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ r2_filter_2(X0, X1))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f213])).
fof(f213, plain, ! [X0] : (! [X1] : ((r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f212])).
fof(f212, plain, ! [X0] : (! [X1] : ((r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f90])).
fof(f90, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', t33_filter_2)).
fof(f10690, plain, (~ v1_filter_0(sK3, k1_lattice2(sK2)) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | (~ spl33_1 | ~ spl33_2 | spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f10672, f541])).
fof(f541, plain, (r2_hidden(sK4, sK3) | ~ spl33_1), inference(avatar_component_clause, [], [f540])).
fof(f540, plain, (spl33_1 <=> r2_hidden(sK4, sK3)), introduced(avatar_definition, [new_symbols(naming, [spl33_1])])).
fof(f10672, plain, (~ r2_hidden(sK4, sK3) | ~ v1_filter_0(sK3, k1_lattice2(sK2)) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | (~ spl33_2 | spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(resolution, [], [f1330, f546])).
fof(f546, plain, (r2_hidden(k7_lattices(sK2, sK4), sK3) | ~ spl33_2), inference(avatar_component_clause, [], [f544])).
fof(f544, plain, (spl33_2 <=> r2_hidden(k7_lattices(sK2, sK4), sK3)), introduced(avatar_definition, [new_symbols(naming, [spl33_2])])).
fof(f1330, plain, (! [X1] : (~ r2_hidden(k7_lattices(sK2, sK4), X1) | ~ r2_hidden(sK4, X1) | ~ v1_filter_0(X1, k1_lattice2(sK2)) | ~ m1_filter_0(X1, k1_lattice2(sK2))) | (spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f1329, f1094])).
fof(f1329, plain, (! [X1] : (~ r2_hidden(k7_lattices(sK2, sK4), X1) | ~ r2_hidden(sK4, X1) | ~ v1_filter_0(X1, k1_lattice2(sK2)) | ~ m1_filter_0(X1, k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | (~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f1328, f1098])).
fof(f1328, plain, (! [X1] : (~ r2_hidden(k7_lattices(sK2, sK4), X1) | ~ r2_hidden(sK4, X1) | ~ v1_filter_0(X1, k1_lattice2(sK2)) | ~ m1_filter_0(X1, k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | (~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f1327, f1190])).
fof(f1190, plain, (v17_lattices(k1_lattice2(sK2)) | ~ spl33_35), inference(avatar_component_clause, [], [f1189])).
fof(f1189, plain, (spl33_35 <=> v17_lattices(k1_lattice2(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl33_35])])).
fof(f1327, plain, (! [X1] : (~ r2_hidden(k7_lattices(sK2, sK4), X1) | ~ r2_hidden(sK4, X1) | ~ v1_filter_0(X1, k1_lattice2(sK2)) | ~ m1_filter_0(X1, k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | ~ spl33_31), inference(subsumption_resolution, [], [f1326, f1102])).
fof(f1326, plain, ! [X1] : (~ r2_hidden(k7_lattices(sK2, sK4), X1) | ~ r2_hidden(sK4, X1) | ~ v1_filter_0(X1, k1_lattice2(sK2)) | ~ m1_filter_0(X1, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))), inference(subsumption_resolution, [], [f1319, f1089])).
fof(f1089, plain, m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))), inference(subsumption_resolution, [], [f1088, f295])).
fof(f1088, plain, (m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1087, f296])).
fof(f1087, plain, (m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1086, f298])).
fof(f1086, plain, (m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1083, f301])).
fof(f301, plain, m1_subset_1(sK4, u1_struct_0(sK2)), inference(cnf_transformation, [], [f232])).
fof(f1083, plain, (m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | ~ m1_subset_1(sK4, u1_struct_0(sK2)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(superposition, [], [f348, f1051])).
fof(f1051, plain, (sK4 = k5_filter_2(sK2, sK4)), inference(subsumption_resolution, [], [f1050, f295])).
fof(f1050, plain, ((sK4 = k5_filter_2(sK2, sK4)) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1049, f296])).
fof(f1049, plain, ((sK4 = k5_filter_2(sK2, sK4)) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1046, f298])).
fof(f1046, plain, ((sK4 = k5_filter_2(sK2, sK4)) | ~ l3_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f341, f301])).
fof(f341, plain, ! [X0, X1] : (~ m1_subset_1(X1, u1_struct_0(X0)) | (k5_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ! [X0] : (! [X1] : ((k5_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f137])).
fof(f137, plain, ! [X0] : (! [X1] : ((k5_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => (k5_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', d4_filter_2)).
fof(f348, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | ~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f146])).
fof(f146, plain, ! [X0, X1] : (m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0))) | (~ m1_subset_1(X1, u1_struct_0(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ! [X0, X1] : ((m1_subset_1(X1, u1_struct_0(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_subset_1(k5_filter_2(X0, X1), u1_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', dt_k5_filter_2)).
fof(f1319, plain, ! [X1] : (~ r2_hidden(k7_lattices(sK2, sK4), X1) | ~ r2_hidden(sK4, X1) | ~ m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | ~ v1_filter_0(X1, k1_lattice2(sK2)) | ~ m1_filter_0(X1, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))), inference(superposition, [], [f533, f1182])).
fof(f1182, plain, (k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), sK4)), inference(forward_demodulation, [], [f1181, f1051])).
fof(f1181, plain, (k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), k5_filter_2(sK2, sK4))), inference(subsumption_resolution, [], [f1180, f295])).
fof(f1180, plain, ((k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), k5_filter_2(sK2, sK4))) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1179, f296])).
fof(f1179, plain, ((k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), k5_filter_2(sK2, sK4))) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1178, f297])).
fof(f297, plain, v17_lattices(sK2), inference(cnf_transformation, [], [f232])).
fof(f1178, plain, ((k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), k5_filter_2(sK2, sK4))) | ~ v17_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f1172, f298])).
fof(f1172, plain, ((k7_lattices(sK2, sK4) = k7_lattices(k1_lattice2(sK2), k5_filter_2(sK2, sK4))) | ~ l3_lattices(sK2) | ~ v17_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f430, f301])).
fof(f430, plain, ! [X0, X1] : (~ m1_subset_1(X1, u1_struct_0(X0)) | (k7_lattices(X0, X1) = k7_lattices(k1_lattice2(X0), k5_filter_2(X0, X1))) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ! [X0] : (! [X1] : ((k7_lattices(X0, X1) = k7_lattices(k1_lattice2(X0), k5_filter_2(X0, X1))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f198])).
fof(f198, plain, ! [X0] : (! [X1] : ((k7_lattices(X0, X1) = k7_lattices(k1_lattice2(X0), k5_filter_2(X0, X1))) | ~ m1_subset_1(X1, u1_struct_0(X0))) | (~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f66])).
fof(f66, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, u1_struct_0(X0)) => (k7_lattices(X0, X1) = k7_lattices(k1_lattice2(X0), k5_filter_2(X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', l70_filter_2)).
fof(f533, plain, ! [X2, X0, X1] : (~ r2_hidden(k7_lattices(X0, X2), X1) | ~ r2_hidden(X2, X1) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ v1_filter_0(X1, X0) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f294])).
fof(f294, plain, ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X2, X1) | r2_hidden(k7_lattices(X0, X2), X1)) & (~ r2_hidden(k7_lattices(X0, X2), X1) | ~ r2_hidden(X2, X1))) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ v1_filter_0(X1, X0) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f218])).
fof(f218, plain, ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X2, X1) <=> ~ r2_hidden(k7_lattices(X0, X2), X1)) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ v1_filter_0(X1, X0) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f217])).
fof(f217, plain, ! [X0] : (! [X1] : ((! [X2] : ((r2_hidden(X2, X1) <=> ~ r2_hidden(k7_lattices(X0, X2), X1)) | ~ m1_subset_1(X2, u1_struct_0(X0))) | ~ v1_filter_0(X1, X0)) | ~ m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f93])).
fof(f93, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_0(X1, X0) => (v1_filter_0(X1, X0) => ! [X2] : (m1_subset_1(X2, u1_struct_0(X0)) => (r2_hidden(X2, X1) <=> ~ r2_hidden(k7_lattices(X0, X2), X1)))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', t59_filter_0)).
fof(f10640, plain, (spl33_1 | spl33_2 | spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35), inference(avatar_contradiction_clause, [], [f10639])).
fof(f10639, plain, ($false | (spl33_1 | spl33_2 | spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f10638, f2110])).
fof(f10638, plain, (~ m1_filter_0(sK3, k1_lattice2(sK2)) | (spl33_1 | spl33_2 | spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f10637, f2108])).
fof(f10637, plain, (~ v1_filter_0(sK3, k1_lattice2(sK2)) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | (spl33_1 | spl33_2 | spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f10605, f542])).
fof(f542, plain, (~ r2_hidden(sK4, sK3) | spl33_1), inference(avatar_component_clause, [], [f540])).
fof(f10605, plain, (r2_hidden(sK4, sK3) | ~ v1_filter_0(sK3, k1_lattice2(sK2)) | ~ m1_filter_0(sK3, k1_lattice2(sK2)) | (spl33_2 | spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(resolution, [], [f1325, f545])).
fof(f545, plain, (~ r2_hidden(k7_lattices(sK2, sK4), sK3) | spl33_2), inference(avatar_component_clause, [], [f544])).
fof(f1325, plain, (! [X0] : (r2_hidden(k7_lattices(sK2, sK4), X0) | r2_hidden(sK4, X0) | ~ v1_filter_0(X0, k1_lattice2(sK2)) | ~ m1_filter_0(X0, k1_lattice2(sK2))) | (spl33_29 | ~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f1324, f1094])).
fof(f1324, plain, (! [X0] : (r2_hidden(k7_lattices(sK2, sK4), X0) | r2_hidden(sK4, X0) | ~ v1_filter_0(X0, k1_lattice2(sK2)) | ~ m1_filter_0(X0, k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | (~ spl33_30 | ~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f1323, f1098])).
fof(f1323, plain, (! [X0] : (r2_hidden(k7_lattices(sK2, sK4), X0) | r2_hidden(sK4, X0) | ~ v1_filter_0(X0, k1_lattice2(sK2)) | ~ m1_filter_0(X0, k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | (~ spl33_31 | ~ spl33_35)), inference(subsumption_resolution, [], [f1322, f1190])).
fof(f1322, plain, (! [X0] : (r2_hidden(k7_lattices(sK2, sK4), X0) | r2_hidden(sK4, X0) | ~ v1_filter_0(X0, k1_lattice2(sK2)) | ~ m1_filter_0(X0, k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))) | ~ spl33_31), inference(subsumption_resolution, [], [f1321, f1102])).
fof(f1321, plain, ! [X0] : (r2_hidden(k7_lattices(sK2, sK4), X0) | r2_hidden(sK4, X0) | ~ v1_filter_0(X0, k1_lattice2(sK2)) | ~ m1_filter_0(X0, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))), inference(subsumption_resolution, [], [f1318, f1089])).
fof(f1318, plain, ! [X0] : (r2_hidden(k7_lattices(sK2, sK4), X0) | r2_hidden(sK4, X0) | ~ m1_subset_1(sK4, u1_struct_0(k1_lattice2(sK2))) | ~ v1_filter_0(X0, k1_lattice2(sK2)) | ~ m1_filter_0(X0, k1_lattice2(sK2)) | ~ l3_lattices(k1_lattice2(sK2)) | ~ v17_lattices(k1_lattice2(sK2)) | ~ v10_lattices(k1_lattice2(sK2)) | v3_struct_0(k1_lattice2(sK2))), inference(superposition, [], [f534, f1182])).
fof(f534, plain, ! [X2, X0, X1] : (r2_hidden(k7_lattices(X0, X2), X1) | r2_hidden(X2, X1) | ~ m1_subset_1(X2, u1_struct_0(X0)) | ~ v1_filter_0(X1, X0) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f294])).
fof(f1231, plain, spl33_35, inference(avatar_contradiction_clause, [], [f1230])).
fof(f1230, plain, ($false | spl33_35), inference(subsumption_resolution, [], [f1229, f942])).
fof(f942, plain, sP0(sK2), inference(subsumption_resolution, [], [f941, f295])).
fof(f941, plain, (sP0(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f940, f296])).
fof(f940, plain, (sP0(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(subsumption_resolution, [], [f931, f297])).
fof(f931, plain, (sP0(sK2) | ~ v17_lattices(sK2) | ~ v10_lattices(sK2) | v3_struct_0(sK2)), inference(resolution, [], [f409, f298])).
fof(f409, plain, ! [X0] : (~ l3_lattices(X0) | sP0(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f224])).
fof(f224, plain, ! [X0] : (sP0(X0) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(definition_folding, [], [f187, e223])).
fof(f223, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(usedef, [], [e223])).
fof(e223, plain, ! [X0] : (sP0(X0) <=> (v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f187, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f186])).
fof(f186, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v17_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f109])).
fof(f109, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), inference(pure_predicate_removal, [], [f60])).
fof(f60, plain, ! [X0] : ((l3_lattices(X0) & v17_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v12_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', fc4_filter_2)).
fof(f1229, plain, (~ sP0(sK2) | spl33_35), inference(resolution, [], [f1191, f408])).
fof(f408, plain, ! [X0] : (v17_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f255])).
fof(f255, plain, ! [X0] : ((v17_lattices(k1_lattice2(X0)) & v16_lattices(k1_lattice2(X0)) & v15_lattices(k1_lattice2(X0)) & v14_lattices(k1_lattice2(X0)) & v13_lattices(k1_lattice2(X0)) & v11_lattices(k1_lattice2(X0)) & v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(nnf_transformation, [], [f223])).
fof(f1191, plain, (~ v17_lattices(k1_lattice2(sK2)) | spl33_35), inference(avatar_component_clause, [], [f1189])).
fof(f1159, plain, ~ spl33_29, inference(avatar_contradiction_clause, [], [f1158])).
fof(f1158, plain, ($false | ~ spl33_29), inference(subsumption_resolution, [], [f1152, f942])).
fof(f1152, plain, (~ sP0(sK2) | ~ spl33_29), inference(resolution, [], [f1095, f394])).
fof(f394, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f255])).
fof(f1095, plain, (v3_struct_0(k1_lattice2(sK2)) | ~ spl33_29), inference(avatar_component_clause, [], [f1093])).
fof(f1147, plain, spl33_31, inference(avatar_contradiction_clause, [], [f1146])).
fof(f1146, plain, ($false | spl33_31), inference(subsumption_resolution, [], [f1145, f298])).
fof(f1145, plain, (~ l3_lattices(sK2) | spl33_31), inference(resolution, [], [f1103, f347])).
fof(f347, plain, ! [X0] : (l3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ! [X0] : ((l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : (l3_lattices(X0) => (l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT323+1.p', dt_k1_lattice2)).
fof(f1103, plain, (~ l3_lattices(k1_lattice2(sK2)) | spl33_31), inference(avatar_component_clause, [], [f1101])).
fof(f1136, plain, spl33_30, inference(avatar_contradiction_clause, [], [f1135])).
fof(f1135, plain, ($false | spl33_30), inference(subsumption_resolution, [], [f1132, f942])).
fof(f1132, plain, (~ sP0(sK2) | spl33_30), inference(resolution, [], [f1099, f402])).
fof(f402, plain, ! [X0] : (v10_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f255])).
fof(f1099, plain, (~ v10_lattices(k1_lattice2(sK2)) | spl33_30), inference(avatar_component_clause, [], [f1097])).
fof(f548, plain, (spl33_1 | ~ spl33_2), inference(avatar_split_clause, [], [f302, f544, f540])).
fof(f302, plain, (~ r2_hidden(k7_lattices(sK2, sK4), sK3) | r2_hidden(sK4, sK3)), inference(cnf_transformation, [], [f232])).
fof(f547, plain, (~ spl33_1 | spl33_2), inference(avatar_split_clause, [], [f303, f544, f540])).
fof(f303, plain, (r2_hidden(k7_lattices(sK2, sK4), sK3) | ~ r2_hidden(sK4, sK3)), inference(cnf_transformation, [], [f232])).