fof(f85064, plain, $false, inference(avatar_sat_refutation, [], [f824, f850, f860, f883, f1206, f1209, f2852, f3256, f3266, f3269, f3283, f3286, f62928, f62943, f78948, f82159, f82179, f84689, f85063])).
fof(f85063, plain, (spl29_11 | ~ spl29_12 | ~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | ~ spl29_2025 | ~ spl29_2027), inference(avatar_contradiction_clause, [], [f85062])).
fof(f85062, plain, ($false | (spl29_11 | ~ spl29_12 | ~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | ~ spl29_2025 | ~ spl29_2027)), inference(subsumption_resolution, [], [f85061, f810])).
fof(f810, plain, (~ v3_struct_0(k1_lattice2(sK1)) | spl29_11), inference(avatar_component_clause, [], [f809])).
fof(f809, plain, (spl29_11 <=> v3_struct_0(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_11])])).
fof(f85061, plain, (v3_struct_0(k1_lattice2(sK1)) | (~ spl29_12 | ~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | ~ spl29_2025 | ~ spl29_2027)), inference(subsumption_resolution, [], [f85060, f814])).
fof(f814, plain, (v10_lattices(k1_lattice2(sK1)) | ~ spl29_12), inference(avatar_component_clause, [], [f813])).
fof(f813, plain, (spl29_12 <=> v10_lattices(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_12])])).
fof(f85060, plain, (~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | ~ spl29_2025 | ~ spl29_2027)), inference(subsumption_resolution, [], [f85059, f818])).
fof(f818, plain, (l3_lattices(k1_lattice2(sK1)) | ~ spl29_13), inference(avatar_component_clause, [], [f817])).
fof(f817, plain, (spl29_13 <=> l3_lattices(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_13])])).
fof(f85059, plain, (~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl29_14 | ~ spl29_119 | spl29_120 | ~ spl29_2025 | ~ spl29_2027)), inference(subsumption_resolution, [], [f85058, f2834])).
fof(f2834, plain, (v13_lattices(k1_lattice2(sK1)) | ~ spl29_119), inference(avatar_component_clause, [], [f2833])).
fof(f2833, plain, (spl29_119 <=> v13_lattices(k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_119])])).
fof(f85058, plain, (~ v13_lattices(k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl29_14 | spl29_120 | ~ spl29_2025 | ~ spl29_2027)), inference(subsumption_resolution, [], [f85057, f1576])).
fof(f1576, plain, (m1_filter_0(sK2, k1_lattice2(sK1)) | ~ spl29_14), inference(backward_demodulation, [], [f823, f1573])).
fof(f1573, plain, (sK2 = k7_filter_2(sK1, sK2)), inference(subsumption_resolution, [], [f1572, f272])).
fof(f272, plain, ~ v3_struct_0(sK1), inference(cnf_transformation, [], [f213])).
fof(f213, plain, ((! [X2] : (~ r2_filter_2(sK1, X2) | ~ r1_tarski(sK2, X2) | ~ m2_filter_2(X2, sK1)) & ~ (u1_struct_0(sK1) = sK2) & m2_filter_2(sK2, sK1)) & v14_lattices(sK1) & l3_lattices(sK1) & v10_lattices(sK1) & ~ v3_struct_0(sK1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2])], [f105, f212, f211])).
fof(f211, plain, (? [X0] : (? [X1] : (! [X2] : (~ r2_filter_2(X0, X2) | ~ r1_tarski(X1, X2) | ~ m2_filter_2(X2, X0)) & ~ (u1_struct_0(X0) = X1) & m2_filter_2(X1, X0)) & v14_lattices(X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (? [X1] : (! [X2] : (~ r2_filter_2(sK1, X2) | ~ r1_tarski(X1, X2) | ~ m2_filter_2(X2, sK1)) & ~ (u1_struct_0(sK1) = X1) & m2_filter_2(X1, sK1)) & v14_lattices(sK1) & l3_lattices(sK1) & v10_lattices(sK1) & ~ v3_struct_0(sK1))), introduced(choice_axiom, [])).
fof(f212, plain, (? [X1] : (! [X2] : (~ r2_filter_2(sK1, X2) | ~ r1_tarski(X1, X2) | ~ m2_filter_2(X2, sK1)) & ~ (u1_struct_0(sK1) = X1) & m2_filter_2(X1, sK1)) => (! [X2] : (~ r2_filter_2(sK1, X2) | ~ r1_tarski(sK2, X2) | ~ m2_filter_2(X2, sK1)) & ~ (u1_struct_0(sK1) = sK2) & m2_filter_2(sK2, sK1))), introduced(choice_axiom, [])).
fof(f105, plain, ? [X0] : (? [X1] : (! [X2] : (~ r2_filter_2(X0, X2) | ~ r1_tarski(X1, X2) | ~ m2_filter_2(X2, X0)) & ~ (u1_struct_0(X0) = X1) & m2_filter_2(X1, X0)) & v14_lattices(X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f104])).
fof(f104, plain, ? [X0] : ((? [X1] : ((! [X2] : ((~ r2_filter_2(X0, X2) | ~ r1_tarski(X1, X2)) | ~ m2_filter_2(X2, X0)) & ~ (u1_struct_0(X0) = X1)) & m2_filter_2(X1, X0)) & v14_lattices(X0)) & (l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v14_lattices(X0) => ! [X1] : (m2_filter_2(X1, X0) => ~ (! [X2] : (m2_filter_2(X2, X0) => ~ (r2_filter_2(X0, X2) & r1_tarski(X1, X2))) & ~ (u1_struct_0(X0) = X1))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v14_lattices(X0) => ! [X1] : (m2_filter_2(X1, X0) => ~ (! [X2] : (m2_filter_2(X2, X0) => ~ (r2_filter_2(X0, X2) & r1_tarski(X1, X2))) & ~ (u1_struct_0(X0) = X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', t34_filter_2)).
fof(f1572, plain, (v3_struct_0(sK1) | (sK2 = k7_filter_2(sK1, sK2))), inference(subsumption_resolution, [], [f1571, f273])).
fof(f273, plain, v10_lattices(sK1), inference(cnf_transformation, [], [f213])).
fof(f1571, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | (sK2 = k7_filter_2(sK1, sK2))), inference(subsumption_resolution, [], [f1565, f274])).
fof(f274, plain, l3_lattices(sK1), inference(cnf_transformation, [], [f213])).
fof(f1565, plain, (~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (sK2 = k7_filter_2(sK1, sK2))), inference(resolution, [], [f774, f708])).
fof(f708, plain, m2_lattice4(sK2, sK1), inference(subsumption_resolution, [], [f707, f272])).
fof(f707, plain, (m2_lattice4(sK2, sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f706, f273])).
fof(f706, plain, (m2_lattice4(sK2, sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f703, f274])).
fof(f703, plain, (m2_lattice4(sK2, sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f319, f276])).
fof(f276, plain, m2_filter_2(sK2, sK1), inference(cnf_transformation, [], [f213])).
fof(f319, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f144])).
fof(f144, plain, ! [X0] : (! [X1] : ((m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (m2_lattice4(X1, X0) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', dt_m2_filter_2)).
fof(f774, plain, ! [X0, X1] : (~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | (k7_filter_2(X0, X1) = X1)), inference(duplicate_literal_removal, [], [f765])).
fof(f765, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f300, f320])).
fof(f320, plain, ! [X0, X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f146])).
fof(f146, plain, ! [X0] : (! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m2_lattice4(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f33])).
fof(f33, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_lattice4(X1, X0) => m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', dt_m2_lattice4)).
fof(f300, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | (k7_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f122])).
fof(f122, plain, ! [X0] : (! [X1] : ((k7_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) => (k7_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', d6_filter_2)).
fof(f823, plain, (m1_filter_0(k7_filter_2(sK1, sK2), k1_lattice2(sK1)) | ~ spl29_14), inference(avatar_component_clause, [], [f821])).
fof(f821, plain, (spl29_14 <=> m1_filter_0(k7_filter_2(sK1, sK2), k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_14])])).
fof(f85057, plain, (~ m1_filter_0(sK2, k1_lattice2(sK1)) | ~ v13_lattices(k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (spl29_120 | ~ spl29_2025 | ~ spl29_2027)), inference(subsumption_resolution, [], [f85056, f2838])).
fof(f2838, plain, (~ (sK2 = u1_struct_0(k1_lattice2(sK1))) | spl29_120), inference(avatar_component_clause, [], [f2837])).
fof(f2837, plain, (spl29_120 <=> (sK2 = u1_struct_0(k1_lattice2(sK1)))), introduced(avatar_definition, [new_symbols(naming, [spl29_120])])).
fof(f85056, plain, ((sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ m1_filter_0(sK2, k1_lattice2(sK1)) | ~ v13_lattices(k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl29_2025 | ~ spl29_2027)), inference(resolution, [], [f84854, f430])).
fof(f430, plain, ! [X0, X1] : (r1_tarski(X1, sK28(X0, X1)) | (u1_struct_0(X0) = X1) | ~ m1_filter_0(X1, X0) | ~ v13_lattices(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f268])).
fof(f268, plain, ! [X0] : (! [X1] : ((v1_filter_0(sK28(X0, X1), X0) & r1_tarski(X1, sK28(X0, X1)) & m1_filter_0(sK28(X0, X1), X0)) | (u1_struct_0(X0) = X1) | ~ m1_filter_0(X1, X0)) | ~ v13_lattices(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK28])], [f196, f267])).
fof(f267, plain, ! [X1, X0] : (? [X2] : (v1_filter_0(X2, X0) & r1_tarski(X1, X2) & m1_filter_0(X2, X0)) => (v1_filter_0(sK28(X0, X1), X0) & r1_tarski(X1, sK28(X0, X1)) & m1_filter_0(sK28(X0, X1), X0))), introduced(choice_axiom, [])).
fof(f196, plain, ! [X0] : (! [X1] : (? [X2] : (v1_filter_0(X2, X0) & r1_tarski(X1, X2) & m1_filter_0(X2, X0)) | (u1_struct_0(X0) = X1) | ~ m1_filter_0(X1, X0)) | ~ v13_lattices(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f195])).
fof(f195, plain, ! [X0] : ((! [X1] : ((? [X2] : ((v1_filter_0(X2, X0) & r1_tarski(X1, X2)) & m1_filter_0(X2, X0)) | (u1_struct_0(X0) = X1)) | ~ m1_filter_0(X1, X0)) | ~ v13_lattices(X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v13_lattices(X0) => ! [X1] : (m1_filter_0(X1, X0) => ~ (! [X2] : (m1_filter_0(X2, X0) => ~ (v1_filter_0(X2, X0) & r1_tarski(X1, X2))) & ~ (u1_struct_0(X0) = X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', t22_filter_0)).
fof(f84854, plain, (~ r1_tarski(sK2, sK28(k1_lattice2(sK1), sK2)) | (~ spl29_2025 | ~ spl29_2027)), inference(subsumption_resolution, [], [f84853, f78922])).
fof(f78922, plain, (m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | ~ spl29_2025), inference(avatar_component_clause, [], [f78921])).
fof(f78921, plain, (spl29_2025 <=> m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1)), introduced(avatar_definition, [new_symbols(naming, [spl29_2025])])).
fof(f84853, plain, (~ r1_tarski(sK2, sK28(k1_lattice2(sK1), sK2)) | ~ m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | ~ spl29_2027), inference(resolution, [], [f78938, f278])).
fof(f278, plain, ! [X2] : (~ r2_filter_2(sK1, X2) | ~ r1_tarski(sK2, X2) | ~ m2_filter_2(X2, sK1)), inference(cnf_transformation, [], [f213])).
fof(f78938, plain, (r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | ~ spl29_2027), inference(avatar_component_clause, [], [f78937])).
fof(f78937, plain, (spl29_2027 <=> r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))), introduced(avatar_definition, [new_symbols(naming, [spl29_2027])])).
fof(f84689, plain, (spl29_11 | ~ spl29_12 | ~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | spl29_2028), inference(avatar_contradiction_clause, [], [f84688])).
fof(f84688, plain, ($false | (spl29_11 | ~ spl29_12 | ~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | spl29_2028)), inference(subsumption_resolution, [], [f84687, f810])).
fof(f84687, plain, (v3_struct_0(k1_lattice2(sK1)) | (~ spl29_12 | ~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | spl29_2028)), inference(subsumption_resolution, [], [f84686, f814])).
fof(f84686, plain, (~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | spl29_2028)), inference(subsumption_resolution, [], [f84685, f818])).
fof(f84685, plain, (~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl29_14 | ~ spl29_119 | spl29_120 | spl29_2028)), inference(subsumption_resolution, [], [f84684, f2834])).
fof(f84684, plain, (~ v13_lattices(k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl29_14 | spl29_120 | spl29_2028)), inference(subsumption_resolution, [], [f84683, f1576])).
fof(f84683, plain, (~ m1_filter_0(sK2, k1_lattice2(sK1)) | ~ v13_lattices(k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (spl29_120 | spl29_2028)), inference(subsumption_resolution, [], [f84682, f2838])).
fof(f84682, plain, ((sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ m1_filter_0(sK2, k1_lattice2(sK1)) | ~ v13_lattices(k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | spl29_2028), inference(resolution, [], [f78942, f431])).
fof(f431, plain, ! [X0, X1] : (v1_filter_0(sK28(X0, X1), X0) | (u1_struct_0(X0) = X1) | ~ m1_filter_0(X1, X0) | ~ v13_lattices(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f268])).
fof(f78942, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | spl29_2028), inference(avatar_component_clause, [], [f78941])).
fof(f78941, plain, (spl29_2028 <=> v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_2028])])).
fof(f82179, plain, (spl29_2025 | ~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119 | ~ spl29_1949), inference(avatar_split_clause, [], [f82178, f78004, f2833, f2829, f1204, f821, f813, f78921])).
fof(f1204, plain, (spl29_63 <=> ! [X16, X15, X17] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15))), introduced(avatar_definition, [new_symbols(naming, [spl29_63])])).
fof(f2829, plain, (spl29_118 <=> (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))), introduced(avatar_definition, [new_symbols(naming, [spl29_118])])).
fof(f78004, plain, (spl29_1949 <=> m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_1949])])).
fof(f82178, plain, (m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119 | ~ spl29_1949)), inference(subsumption_resolution, [], [f82177, f272])).
fof(f82177, plain, (m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | v3_struct_0(sK1) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119 | ~ spl29_1949)), inference(subsumption_resolution, [], [f82176, f273])).
fof(f82176, plain, (m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119 | ~ spl29_1949)), inference(subsumption_resolution, [], [f82175, f274])).
fof(f82175, plain, (m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119 | ~ spl29_1949)), inference(subsumption_resolution, [], [f78907, f78005])).
fof(f78005, plain, (m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | ~ spl29_1949), inference(avatar_component_clause, [], [f78004])).
fof(f78907, plain, (m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | ~ m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119)), inference(superposition, [], [f304, f63071])).
fof(f63071, plain, ((sK28(k1_lattice2(sK1), sK2) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119)), inference(forward_demodulation, [], [f2831, f15366])).
fof(f15366, plain, ((sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_119)), inference(subsumption_resolution, [], [f15365, f277])).
fof(f277, plain, ~ (u1_struct_0(sK1) = sK2), inference(cnf_transformation, [], [f213])).
fof(f15365, plain, ((u1_struct_0(sK1) = sK2) | (sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_119)), inference(forward_demodulation, [], [f15364, f4966])).
fof(f4966, plain, ((u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | ~ spl29_63), inference(trivial_inequality_removal, [], [f4949])).
fof(f4949, plain, (~ (k1_lattice2(sK1) = k1_lattice2(sK1)) | (u1_struct_0(sK1) = u1_struct_0(k1_lattice2(sK1))) | ~ spl29_63), inference(superposition, [], [f1205, f1595])).
fof(f1595, plain, (k1_lattice2(sK1) = g3_lattices(u1_struct_0(k1_lattice2(sK1)), u2_lattices(k1_lattice2(sK1)), u1_lattices(k1_lattice2(sK1)))), inference(resolution, [], [f731, f274])).
fof(f731, plain, ! [X1] : (~ l3_lattices(X1) | (k1_lattice2(X1) = g3_lattices(u1_struct_0(k1_lattice2(X1)), u2_lattices(k1_lattice2(X1)), u1_lattices(k1_lattice2(X1))))), inference(subsumption_resolution, [], [f726, f307])).
fof(f307, plain, ! [X0] : (l3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ! [X0] : ((l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0] : (l3_lattices(X0) => (l3_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', dt_k1_lattice2)).
fof(f726, plain, ! [X1] : ((k1_lattice2(X1) = g3_lattices(u1_struct_0(k1_lattice2(X1)), u2_lattices(k1_lattice2(X1)), u1_lattices(k1_lattice2(X1)))) | ~ l3_lattices(k1_lattice2(X1)) | ~ l3_lattices(X1)), inference(resolution, [], [f279, f306])).
fof(f306, plain, ! [X0] : (v3_lattices(k1_lattice2(X0)) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f132])).
fof(f279, plain, ! [X0] : (~ v3_lattices(X0) | (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ l3_lattices(X0)), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ! [X0] : ((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0) | ~ l3_lattices(X0)), inference(flattening, [], [f106])).
fof(f106, plain, ! [X0] : (((g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0) | ~ v3_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (l3_lattices(X0) => (v3_lattices(X0) => (g3_lattices(u1_struct_0(X0), u2_lattices(X0), u1_lattices(X0)) = X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', abstractness_v3_lattices)).
fof(f1205, plain, (! [X17, X15, X16] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15)) | ~ spl29_63), inference(avatar_component_clause, [], [f1204])).
fof(f15364, plain, ((sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_119)), inference(subsumption_resolution, [], [f15363, f2834])).
fof(f15363, plain, ((sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ v13_lattices(k1_lattice2(sK1)) | (~ spl29_12 | ~ spl29_14)), inference(subsumption_resolution, [], [f15362, f274])).
fof(f15362, plain, ((sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(sK1) | ~ v13_lattices(k1_lattice2(sK1)) | (~ spl29_12 | ~ spl29_14)), inference(subsumption_resolution, [], [f15361, f814])).
fof(f15361, plain, ((sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ v10_lattices(k1_lattice2(sK1)) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(sK1) | ~ v13_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(subsumption_resolution, [], [f15360, f272])).
fof(f15360, plain, (v3_struct_0(sK1) | (sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ v10_lattices(k1_lattice2(sK1)) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(sK1) | ~ v13_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(subsumption_resolution, [], [f15347, f273])).
fof(f15347, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | (sK28(k1_lattice2(sK1), sK2) = k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ v10_lattices(k1_lattice2(sK1)) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(sK1) | ~ v13_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(resolution, [], [f2089, f1576])).
fof(f2089, plain, ! [X4, X3] : (~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v10_lattices(X3) | v3_struct_0(X3) | (sK28(k1_lattice2(X3), X4) = k8_filter_2(X3, sK28(k1_lattice2(X3), X4))) | ~ v10_lattices(k1_lattice2(X3)) | (u1_struct_0(k1_lattice2(X3)) = X4) | ~ l3_lattices(X3) | ~ v13_lattices(k1_lattice2(X3))), inference(subsumption_resolution, [], [f2088, f340])).
fof(f340, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ! [X0] : ((v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f161])).
fof(f161, plain, ! [X0] : ((v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f50])).
fof(f50, plain, ! [X0] : ((l3_lattices(X0) & ~ v3_struct_0(X0)) => (v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', fc1_lattice2)).
fof(f2088, plain, ! [X4, X3] : (~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | (sK28(k1_lattice2(X3), X4) = k8_filter_2(X3, sK28(k1_lattice2(X3), X4))) | ~ v10_lattices(k1_lattice2(X3)) | (u1_struct_0(k1_lattice2(X3)) = X4) | ~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v13_lattices(k1_lattice2(X3)) | v3_struct_0(k1_lattice2(X3))), inference(subsumption_resolution, [], [f2083, f307])).
fof(f2083, plain, ! [X4, X3] : (~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | (sK28(k1_lattice2(X3), X4) = k8_filter_2(X3, sK28(k1_lattice2(X3), X4))) | ~ v10_lattices(k1_lattice2(X3)) | (u1_struct_0(k1_lattice2(X3)) = X4) | ~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v13_lattices(k1_lattice2(X3)) | ~ l3_lattices(k1_lattice2(X3)) | v3_struct_0(k1_lattice2(X3))), inference(duplicate_literal_removal, [], [f2080])).
fof(f2080, plain, ! [X4, X3] : (~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | (sK28(k1_lattice2(X3), X4) = k8_filter_2(X3, sK28(k1_lattice2(X3), X4))) | ~ v10_lattices(k1_lattice2(X3)) | (u1_struct_0(k1_lattice2(X3)) = X4) | ~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v13_lattices(k1_lattice2(X3)) | ~ l3_lattices(k1_lattice2(X3)) | ~ v10_lattices(k1_lattice2(X3)) | v3_struct_0(k1_lattice2(X3))), inference(resolution, [], [f802, f429])).
fof(f429, plain, ! [X0, X1] : (m1_filter_0(sK28(X0, X1), X0) | (u1_struct_0(X0) = X1) | ~ m1_filter_0(X1, X0) | ~ v13_lattices(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f268])).
fof(f802, plain, ! [X2, X3] : (~ m1_filter_0(X3, k1_lattice2(X2)) | ~ l3_lattices(X2) | ~ v10_lattices(X2) | v3_struct_0(X2) | (k8_filter_2(X2, X3) = X3) | ~ v10_lattices(k1_lattice2(X2))), inference(subsumption_resolution, [], [f801, f340])).
fof(f801, plain, ! [X2, X3] : ((k8_filter_2(X2, X3) = X3) | ~ l3_lattices(X2) | ~ v10_lattices(X2) | v3_struct_0(X2) | ~ m1_filter_0(X3, k1_lattice2(X2)) | ~ v10_lattices(k1_lattice2(X2)) | v3_struct_0(k1_lattice2(X2))), inference(subsumption_resolution, [], [f793, f307])).
fof(f793, plain, ! [X2, X3] : ((k8_filter_2(X2, X3) = X3) | ~ l3_lattices(X2) | ~ v10_lattices(X2) | v3_struct_0(X2) | ~ m1_filter_0(X3, k1_lattice2(X2)) | ~ l3_lattices(k1_lattice2(X2)) | ~ v10_lattices(k1_lattice2(X2)) | v3_struct_0(k1_lattice2(X2))), inference(resolution, [], [f301, f315])).
fof(f315, plain, ! [X0, X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ! [X0] : (! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) | ~ m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f140])).
fof(f140, plain, ! [X0] : (! [X1] : ((m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)) | ~ m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f28])).
fof(f28, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_0(X1, X0) => (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) & ~ v1_xboole_0(X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', dt_m1_filter_0)).
fof(f301, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) | (k8_filter_2(X0, X1) = X1) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ! [X0] : (! [X1] : ((k8_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0))))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f124])).
fof(f124, plain, ! [X0] : (! [X1] : ((k8_filter_2(X0, X1) = X1) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0))))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(k1_lattice2(X0)))) => (k8_filter_2(X0, X1) = X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', d7_filter_2)).
fof(f2831, plain, ((k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | ~ spl29_118), inference(avatar_component_clause, [], [f2829])).
fof(f304, plain, ! [X0, X1] : (m2_filter_2(k14_filter_2(X0, X1), X0) | ~ m1_filter_0(X1, k1_lattice2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ! [X0, X1] : (m2_filter_2(k14_filter_2(X0, X1), X0) | ~ m1_filter_0(X1, k1_lattice2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f128])).
fof(f128, plain, ! [X0, X1] : (m2_filter_2(k14_filter_2(X0, X1), X0) | (~ m1_filter_0(X1, k1_lattice2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0, X1] : ((m1_filter_0(X1, k1_lattice2(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m2_filter_2(k14_filter_2(X0, X1), X0)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', dt_k14_filter_2)).
fof(f82159, plain, (spl29_11 | ~ spl29_12 | ~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | spl29_1949), inference(avatar_contradiction_clause, [], [f82158])).
fof(f82158, plain, ($false | (spl29_11 | ~ spl29_12 | ~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | spl29_1949)), inference(subsumption_resolution, [], [f82157, f810])).
fof(f82157, plain, (v3_struct_0(k1_lattice2(sK1)) | (~ spl29_12 | ~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | spl29_1949)), inference(subsumption_resolution, [], [f82156, f814])).
fof(f82156, plain, (~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl29_13 | ~ spl29_14 | ~ spl29_119 | spl29_120 | spl29_1949)), inference(subsumption_resolution, [], [f82155, f818])).
fof(f82155, plain, (~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl29_14 | ~ spl29_119 | spl29_120 | spl29_1949)), inference(subsumption_resolution, [], [f82154, f2834])).
fof(f82154, plain, (~ v13_lattices(k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (~ spl29_14 | spl29_120 | spl29_1949)), inference(subsumption_resolution, [], [f82153, f1576])).
fof(f82153, plain, (~ m1_filter_0(sK2, k1_lattice2(sK1)) | ~ v13_lattices(k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | (spl29_120 | spl29_1949)), inference(subsumption_resolution, [], [f82152, f2838])).
fof(f82152, plain, ((sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ m1_filter_0(sK2, k1_lattice2(sK1)) | ~ v13_lattices(k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1)) | spl29_1949), inference(resolution, [], [f78006, f429])).
fof(f78006, plain, (~ m1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | spl29_1949), inference(avatar_component_clause, [], [f78004])).
fof(f78948, plain, (~ spl29_2025 | spl29_2027 | ~ spl29_2028 | ~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119), inference(avatar_split_clause, [], [f78947, f2833, f2829, f1204, f821, f813, f78941, f78937, f78921])).
fof(f78947, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | ~ m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119)), inference(subsumption_resolution, [], [f78946, f272])).
fof(f78946, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | ~ m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | v3_struct_0(sK1) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119)), inference(subsumption_resolution, [], [f78945, f273])).
fof(f78945, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | ~ m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119)), inference(subsumption_resolution, [], [f78911, f274])).
fof(f78911, plain, (~ v1_filter_0(sK28(k1_lattice2(sK1), sK2), k1_lattice2(sK1)) | r2_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) | ~ m2_filter_2(sK28(k1_lattice2(sK1), sK2), sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119)), inference(superposition, [], [f434, f63072])).
fof(f63072, plain, ((sK28(k1_lattice2(sK1), sK2) = k15_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_118 | ~ spl29_119)), inference(backward_demodulation, [], [f62937, f63071])).
fof(f62937, plain, ((k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_119)), inference(backward_demodulation, [], [f62934, f62936])).
fof(f62936, plain, ((k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_119)), inference(subsumption_resolution, [], [f62935, f277])).
fof(f62935, plain, ((u1_struct_0(sK1) = sK2) | (k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_119)), inference(forward_demodulation, [], [f55894, f4966])).
fof(f55894, plain, ((k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_119)), inference(subsumption_resolution, [], [f55893, f814])).
fof(f55893, plain, ((k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ v10_lattices(k1_lattice2(sK1)) | (~ spl29_14 | ~ spl29_119)), inference(subsumption_resolution, [], [f55892, f2834])).
fof(f55892, plain, ((k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(subsumption_resolution, [], [f55891, f273])).
fof(f55891, plain, ((k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ v10_lattices(sK1) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(subsumption_resolution, [], [f55890, f274])).
fof(f55890, plain, (~ l3_lattices(sK1) | (k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ v10_lattices(sK1) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(subsumption_resolution, [], [f55873, f272])).
fof(f55873, plain, (v3_struct_0(sK1) | ~ l3_lattices(sK1) | (k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ v10_lattices(sK1) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(resolution, [], [f9185, f1576])).
fof(f9185, plain, ! [X4, X3] : (~ m1_filter_0(X4, k1_lattice2(X3)) | v3_struct_0(X3) | ~ l3_lattices(X3) | (k14_filter_2(X3, sK28(k1_lattice2(X3), X4)) = k7_filter_2(X3, k14_filter_2(X3, sK28(k1_lattice2(X3), X4)))) | (u1_struct_0(k1_lattice2(X3)) = X4) | ~ v10_lattices(X3) | ~ v13_lattices(k1_lattice2(X3)) | ~ v10_lattices(k1_lattice2(X3))), inference(subsumption_resolution, [], [f9184, f340])).
fof(f9184, plain, ! [X4, X3] : (~ v10_lattices(X3) | v3_struct_0(X3) | ~ l3_lattices(X3) | (k14_filter_2(X3, sK28(k1_lattice2(X3), X4)) = k7_filter_2(X3, k14_filter_2(X3, sK28(k1_lattice2(X3), X4)))) | (u1_struct_0(k1_lattice2(X3)) = X4) | ~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v13_lattices(k1_lattice2(X3)) | ~ v10_lattices(k1_lattice2(X3)) | v3_struct_0(k1_lattice2(X3))), inference(subsumption_resolution, [], [f9159, f307])).
fof(f9159, plain, ! [X4, X3] : (~ v10_lattices(X3) | v3_struct_0(X3) | ~ l3_lattices(X3) | (k14_filter_2(X3, sK28(k1_lattice2(X3), X4)) = k7_filter_2(X3, k14_filter_2(X3, sK28(k1_lattice2(X3), X4)))) | (u1_struct_0(k1_lattice2(X3)) = X4) | ~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v13_lattices(k1_lattice2(X3)) | ~ l3_lattices(k1_lattice2(X3)) | ~ v10_lattices(k1_lattice2(X3)) | v3_struct_0(k1_lattice2(X3))), inference(resolution, [], [f1649, f429])).
fof(f1649, plain, ! [X4, X3] : (~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v10_lattices(X3) | v3_struct_0(X3) | ~ l3_lattices(X3) | (k14_filter_2(X3, X4) = k7_filter_2(X3, k14_filter_2(X3, X4)))), inference(duplicate_literal_removal, [], [f1648])).
fof(f1648, plain, ! [X4, X3] : (~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | ~ m1_filter_0(X4, k1_lattice2(X3)) | ~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | (k14_filter_2(X3, X4) = k7_filter_2(X3, k14_filter_2(X3, X4)))), inference(resolution, [], [f754, f774])).
fof(f754, plain, ! [X0, X1] : (m2_lattice4(k14_filter_2(X1, X0), X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ m1_filter_0(X0, k1_lattice2(X1))), inference(duplicate_literal_removal, [], [f751])).
fof(f751, plain, ! [X0, X1] : (~ m1_filter_0(X0, k1_lattice2(X1)) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | m2_lattice4(k14_filter_2(X1, X0), X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1)), inference(resolution, [], [f304, f319])).
fof(f62934, plain, ((k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_119)), inference(subsumption_resolution, [], [f62933, f277])).
fof(f62933, plain, ((u1_struct_0(sK1) = sK2) | (k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_63 | ~ spl29_119)), inference(forward_demodulation, [], [f16511, f4966])).
fof(f16511, plain, ((k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | (~ spl29_12 | ~ spl29_14 | ~ spl29_119)), inference(subsumption_resolution, [], [f16510, f814])).
fof(f16510, plain, ((k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ v10_lattices(k1_lattice2(sK1)) | (~ spl29_14 | ~ spl29_119)), inference(subsumption_resolution, [], [f16509, f2834])).
fof(f16509, plain, ((k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(subsumption_resolution, [], [f16508, f274])).
fof(f16508, plain, ((k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(sK1) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(subsumption_resolution, [], [f16507, f272])).
fof(f16507, plain, (v3_struct_0(sK1) | (k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(sK1) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(subsumption_resolution, [], [f16494, f273])).
fof(f16494, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | (k7_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) = k15_filter_2(sK1, k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)))) | (sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ l3_lattices(sK1) | ~ v13_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | ~ spl29_14), inference(resolution, [], [f2471, f1576])).
fof(f2471, plain, ! [X4, X3] : (~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v10_lattices(X3) | v3_struct_0(X3) | (k7_filter_2(X3, k14_filter_2(X3, sK28(k1_lattice2(X3), X4))) = k15_filter_2(X3, k14_filter_2(X3, sK28(k1_lattice2(X3), X4)))) | (u1_struct_0(k1_lattice2(X3)) = X4) | ~ l3_lattices(X3) | ~ v13_lattices(k1_lattice2(X3)) | ~ v10_lattices(k1_lattice2(X3))), inference(subsumption_resolution, [], [f2470, f340])).
fof(f2470, plain, ! [X4, X3] : (~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | (k7_filter_2(X3, k14_filter_2(X3, sK28(k1_lattice2(X3), X4))) = k15_filter_2(X3, k14_filter_2(X3, sK28(k1_lattice2(X3), X4)))) | (u1_struct_0(k1_lattice2(X3)) = X4) | ~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v13_lattices(k1_lattice2(X3)) | ~ v10_lattices(k1_lattice2(X3)) | v3_struct_0(k1_lattice2(X3))), inference(subsumption_resolution, [], [f2460, f307])).
fof(f2460, plain, ! [X4, X3] : (~ l3_lattices(X3) | ~ v10_lattices(X3) | v3_struct_0(X3) | (k7_filter_2(X3, k14_filter_2(X3, sK28(k1_lattice2(X3), X4))) = k15_filter_2(X3, k14_filter_2(X3, sK28(k1_lattice2(X3), X4)))) | (u1_struct_0(k1_lattice2(X3)) = X4) | ~ m1_filter_0(X4, k1_lattice2(X3)) | ~ v13_lattices(k1_lattice2(X3)) | ~ l3_lattices(k1_lattice2(X3)) | ~ v10_lattices(k1_lattice2(X3)) | v3_struct_0(k1_lattice2(X3))), inference(resolution, [], [f782, f429])).
fof(f782, plain, ! [X2, X1] : (~ m1_filter_0(X2, k1_lattice2(X1)) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | (k7_filter_2(X1, k14_filter_2(X1, X2)) = k15_filter_2(X1, k14_filter_2(X1, X2)))), inference(duplicate_literal_removal, [], [f781])).
fof(f781, plain, ! [X2, X1] : ((k7_filter_2(X1, k14_filter_2(X1, X2)) = k15_filter_2(X1, k14_filter_2(X1, X2))) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ m1_filter_0(X2, k1_lattice2(X1)) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1)), inference(resolution, [], [f422, f304])).
fof(f422, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | (k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f191])).
fof(f191, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f190])).
fof(f190, plain, ! [X0, X1] : ((k7_filter_2(X0, X1) = k15_filter_2(X0, X1)) | (~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f78])).
fof(f78, plain, ! [X0, X1] : ((m2_filter_2(X1, X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k7_filter_2(X0, X1) = k15_filter_2(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', redefinition_k15_filter_2)).
fof(f434, plain, ! [X0, X1] : (~ v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | r2_filter_2(X0, X1) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f269])).
fof(f269, plain, ! [X0] : (! [X1] : (((r2_filter_2(X0, X1) | ~ v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) & (v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ r2_filter_2(X0, X1))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f200])).
fof(f200, plain, ! [X0] : (! [X1] : ((r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f199])).
fof(f199, plain, ! [X0] : (! [X1] : ((r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))) | ~ m2_filter_2(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f85])).
fof(f85, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m2_filter_2(X1, X0) => (r2_filter_2(X0, X1) <=> v1_filter_0(k15_filter_2(X0, X1), k1_lattice2(X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', t33_filter_2)).
fof(f62943, plain, (spl29_118 | ~ spl29_12 | ~ spl29_63), inference(avatar_split_clause, [], [f62942, f1204, f813, f2829])).
fof(f62942, plain, ((k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_12 | ~ spl29_63)), inference(subsumption_resolution, [], [f62941, f277])).
fof(f62941, plain, ((u1_struct_0(sK1) = sK2) | (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_12 | ~ spl29_63)), inference(forward_demodulation, [], [f62940, f1574])).
fof(f1574, plain, (sK2 = k15_filter_2(sK1, sK2)), inference(backward_demodulation, [], [f786, f1573])).
fof(f786, plain, (k7_filter_2(sK1, sK2) = k15_filter_2(sK1, sK2)), inference(subsumption_resolution, [], [f785, f272])).
fof(f785, plain, ((k7_filter_2(sK1, sK2) = k15_filter_2(sK1, sK2)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f784, f273])).
fof(f784, plain, ((k7_filter_2(sK1, sK2) = k15_filter_2(sK1, sK2)) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f779, f274])).
fof(f779, plain, ((k7_filter_2(sK1, sK2) = k15_filter_2(sK1, sK2)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f422, f276])).
fof(f62940, plain, ((u1_struct_0(sK1) = k15_filter_2(sK1, sK2)) | (k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (~ spl29_12 | ~ spl29_63)), inference(forward_demodulation, [], [f58191, f4966])).
fof(f58191, plain, ((k8_filter_2(sK1, sK28(k1_lattice2(sK1), sK2)) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), sK2))) | (k15_filter_2(sK1, sK2) = u1_struct_0(k1_lattice2(sK1))) | ~ spl29_12), inference(forward_demodulation, [], [f58190, f1574])).
fof(f58190, plain, ((k8_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2))) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2)))) | (k15_filter_2(sK1, sK2) = u1_struct_0(k1_lattice2(sK1))) | ~ spl29_12), inference(subsumption_resolution, [], [f58189, f275])).
fof(f275, plain, v14_lattices(sK1), inference(cnf_transformation, [], [f213])).
fof(f58189, plain, ((k8_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2))) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2)))) | (k15_filter_2(sK1, sK2) = u1_struct_0(k1_lattice2(sK1))) | ~ v14_lattices(sK1) | ~ spl29_12), inference(subsumption_resolution, [], [f58188, f272])).
fof(f58188, plain, ((k8_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2))) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2)))) | v3_struct_0(sK1) | (k15_filter_2(sK1, sK2) = u1_struct_0(k1_lattice2(sK1))) | ~ v14_lattices(sK1) | ~ spl29_12), inference(subsumption_resolution, [], [f58187, f273])).
fof(f58187, plain, ((k8_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2))) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2)))) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (k15_filter_2(sK1, sK2) = u1_struct_0(k1_lattice2(sK1))) | ~ v14_lattices(sK1) | ~ spl29_12), inference(subsumption_resolution, [], [f58186, f274])).
fof(f58186, plain, ((k8_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2))) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2)))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (k15_filter_2(sK1, sK2) = u1_struct_0(k1_lattice2(sK1))) | ~ v14_lattices(sK1) | ~ spl29_12), inference(subsumption_resolution, [], [f58156, f814])).
fof(f58156, plain, (~ v10_lattices(k1_lattice2(sK1)) | (k8_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2))) = k14_filter_2(sK1, sK28(k1_lattice2(sK1), k15_filter_2(sK1, sK2)))) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | (k15_filter_2(sK1, sK2) = u1_struct_0(k1_lattice2(sK1))) | ~ v14_lattices(sK1)), inference(resolution, [], [f17419, f276])).
fof(f17419, plain, ! [X0, X1] : (~ m2_filter_2(X1, X0) | ~ v10_lattices(k1_lattice2(X0)) | (k8_filter_2(X0, sK28(k1_lattice2(X0), k15_filter_2(X0, X1))) = k14_filter_2(X0, sK28(k1_lattice2(X0), k15_filter_2(X0, X1)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | (u1_struct_0(k1_lattice2(X0)) = k15_filter_2(X0, X1)) | ~ v14_lattices(X0)), inference(duplicate_literal_removal, [], [f17412])).
fof(f17412, plain, ! [X0, X1] : ((u1_struct_0(k1_lattice2(X0)) = k15_filter_2(X0, X1)) | ~ v10_lattices(k1_lattice2(X0)) | (k8_filter_2(X0, sK28(k1_lattice2(X0), k15_filter_2(X0, X1))) = k14_filter_2(X0, sK28(k1_lattice2(X0), k15_filter_2(X0, X1)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_filter_2(X1, X0) | ~ v14_lattices(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(resolution, [], [f2823, f439])).
fof(f439, plain, ! [X0] : (v13_lattices(k1_lattice2(X0)) | ~ v14_lattices(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ! [X0] : (((v14_lattices(X0) | ~ v13_lattices(k1_lattice2(X0))) & (v13_lattices(k1_lattice2(X0)) | ~ v14_lattices(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f205])).
fof(f205, plain, ! [X0] : ((v14_lattices(X0) <=> v13_lattices(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f204])).
fof(f204, plain, ! [X0] : ((v14_lattices(X0) <=> v13_lattices(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f89])).
fof(f89, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v14_lattices(X0) <=> v13_lattices(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', t64_lattice2)).
fof(f2823, plain, ! [X0, X1] : (~ v13_lattices(k1_lattice2(X0)) | (u1_struct_0(k1_lattice2(X0)) = k15_filter_2(X0, X1)) | ~ v10_lattices(k1_lattice2(X0)) | (k8_filter_2(X0, sK28(k1_lattice2(X0), k15_filter_2(X0, X1))) = k14_filter_2(X0, sK28(k1_lattice2(X0), k15_filter_2(X0, X1)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_filter_2(X1, X0)), inference(duplicate_literal_removal, [], [f2816])).
fof(f2816, plain, ! [X0, X1] : ((u1_struct_0(k1_lattice2(X0)) = k15_filter_2(X0, X1)) | ~ v13_lattices(k1_lattice2(X0)) | ~ v10_lattices(k1_lattice2(X0)) | (k8_filter_2(X0, sK28(k1_lattice2(X0), k15_filter_2(X0, X1))) = k14_filter_2(X0, sK28(k1_lattice2(X0), k15_filter_2(X0, X1)))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0) | ~ m2_filter_2(X1, X0) | ~ v10_lattices(k1_lattice2(X0))), inference(resolution, [], [f968, f759])).
fof(f759, plain, ! [X0, X1] : (m1_filter_0(k15_filter_2(X1, X0), k1_lattice2(X1)) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | ~ m2_filter_2(X0, X1) | ~ v10_lattices(k1_lattice2(X1))), inference(subsumption_resolution, [], [f758, f340])).
fof(f758, plain, ! [X0, X1] : (~ m2_filter_2(X0, X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | m1_filter_0(k15_filter_2(X1, X0), k1_lattice2(X1)) | ~ v10_lattices(k1_lattice2(X1)) | v3_struct_0(k1_lattice2(X1))), inference(subsumption_resolution, [], [f755, f307])).
fof(f755, plain, ! [X0, X1] : (~ m2_filter_2(X0, X1) | ~ l3_lattices(X1) | ~ v10_lattices(X1) | v3_struct_0(X1) | m1_filter_0(k15_filter_2(X1, X0), k1_lattice2(X1)) | ~ l3_lattices(k1_lattice2(X1)) | ~ v10_lattices(k1_lattice2(X1)) | v3_struct_0(k1_lattice2(X1))), inference(resolution, [], [f305, f423])).
fof(f423, plain, ! [X0, X1] : (~ m1_filter_2(X1, X0) | m1_filter_0(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f265])).
fof(f265, plain, ! [X0] : (! [X1] : ((m1_filter_2(X1, X0) | ~ m1_filter_0(X1, X0)) & (m1_filter_0(X1, X0) | ~ m1_filter_2(X1, X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f193])).
fof(f193, plain, ! [X0] : (! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f192])).
fof(f192, plain, ! [X0] : (! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0)) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f79])).
fof(f79, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_filter_2(X1, X0) <=> m1_filter_0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', redefinition_m1_filter_2)).
fof(f305, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | ~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f130])).
fof(f130, plain, ! [X0, X1] : (m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0)) | (~ m2_filter_2(X1, X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0, X1] : ((m2_filter_2(X1, X0) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => m1_filter_2(k15_filter_2(X0, X1), k1_lattice2(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', dt_k15_filter_2)).
fof(f968, plain, ! [X4, X5] : (~ m1_filter_0(X5, k1_lattice2(X4)) | (u1_struct_0(k1_lattice2(X4)) = X5) | ~ v13_lattices(k1_lattice2(X4)) | ~ v10_lattices(k1_lattice2(X4)) | (k8_filter_2(X4, sK28(k1_lattice2(X4), X5)) = k14_filter_2(X4, sK28(k1_lattice2(X4), X5))) | ~ l3_lattices(X4) | ~ v10_lattices(X4) | v3_struct_0(X4)), inference(subsumption_resolution, [], [f967, f340])).
fof(f967, plain, ! [X4, X5] : ((u1_struct_0(k1_lattice2(X4)) = X5) | ~ m1_filter_0(X5, k1_lattice2(X4)) | ~ v13_lattices(k1_lattice2(X4)) | ~ v10_lattices(k1_lattice2(X4)) | v3_struct_0(k1_lattice2(X4)) | (k8_filter_2(X4, sK28(k1_lattice2(X4), X5)) = k14_filter_2(X4, sK28(k1_lattice2(X4), X5))) | ~ l3_lattices(X4) | ~ v10_lattices(X4) | v3_struct_0(X4)), inference(subsumption_resolution, [], [f964, f307])).
fof(f964, plain, ! [X4, X5] : ((u1_struct_0(k1_lattice2(X4)) = X5) | ~ m1_filter_0(X5, k1_lattice2(X4)) | ~ v13_lattices(k1_lattice2(X4)) | ~ l3_lattices(k1_lattice2(X4)) | ~ v10_lattices(k1_lattice2(X4)) | v3_struct_0(k1_lattice2(X4)) | (k8_filter_2(X4, sK28(k1_lattice2(X4), X5)) = k14_filter_2(X4, sK28(k1_lattice2(X4), X5))) | ~ l3_lattices(X4) | ~ v10_lattices(X4) | v3_struct_0(X4)), inference(resolution, [], [f429, f421])).
fof(f421, plain, ! [X0, X1] : (~ m1_filter_0(X1, k1_lattice2(X0)) | (k8_filter_2(X0, X1) = k14_filter_2(X0, X1)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f189])).
fof(f189, plain, ! [X0, X1] : ((k8_filter_2(X0, X1) = k14_filter_2(X0, X1)) | ~ m1_filter_0(X1, k1_lattice2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f188])).
fof(f188, plain, ! [X0, X1] : ((k8_filter_2(X0, X1) = k14_filter_2(X0, X1)) | (~ m1_filter_0(X1, k1_lattice2(X0)) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f77])).
fof(f77, plain, ! [X0, X1] : ((m1_filter_0(X1, k1_lattice2(X0)) & l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (k8_filter_2(X0, X1) = k14_filter_2(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', redefinition_k14_filter_2)).
fof(f62928, plain, (~ spl29_63 | ~ spl29_120), inference(avatar_contradiction_clause, [], [f62927])).
fof(f62927, plain, ($false | (~ spl29_63 | ~ spl29_120)), inference(subsumption_resolution, [], [f62926, f277])).
fof(f62926, plain, ((u1_struct_0(sK1) = sK2) | (~ spl29_63 | ~ spl29_120)), inference(backward_demodulation, [], [f4966, f2839])).
fof(f2839, plain, ((sK2 = u1_struct_0(k1_lattice2(sK1))) | ~ spl29_120), inference(avatar_component_clause, [], [f2837])).
fof(f3286, plain, spl29_60, inference(avatar_contradiction_clause, [], [f3285])).
fof(f3285, plain, ($false | spl29_60), inference(subsumption_resolution, [], [f3284, f457])).
fof(f457, plain, l2_lattices(sK1), inference(resolution, [], [f313, f274])).
fof(f313, plain, ! [X0] : (~ l3_lattices(X0) | l2_lattices(X0)), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ! [X0] : ((l2_lattices(X0) & l1_lattices(X0)) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ! [X0] : (l3_lattices(X0) => (l2_lattices(X0) & l1_lattices(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', dt_l3_lattices)).
fof(f3284, plain, (~ l2_lattices(sK1) | spl29_60), inference(resolution, [], [f1194, f717])).
fof(f717, plain, ! [X1] : (m1_relset_1(u2_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l2_lattices(X1)), inference(resolution, [], [f327, f425])).
fof(f425, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f266])).
fof(f266, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f80])).
fof(f80, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', redefinition_m2_relset_1)).
fof(f327, plain, ! [X0] : (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ! [X0] : ((m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0))) | ~ l2_lattices(X0)), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ! [X0] : (l2_lattices(X0) => (m2_relset_1(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', dt_u2_lattices)).
fof(f1194, plain, (~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl29_60), inference(avatar_component_clause, [], [f1192])).
fof(f1192, plain, (spl29_60 <=> m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_60])])).
fof(f3283, plain, spl29_59, inference(avatar_contradiction_clause, [], [f3282])).
fof(f3282, plain, ($false | spl29_59), inference(subsumption_resolution, [], [f3281, f457])).
fof(f3281, plain, (~ l2_lattices(sK1) | spl29_59), inference(resolution, [], [f1190, f326])).
fof(f326, plain, ! [X0] : (v1_funct_2(u2_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f150])).
fof(f1190, plain, (~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl29_59), inference(avatar_component_clause, [], [f1188])).
fof(f1188, plain, (spl29_59 <=> v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_59])])).
fof(f3269, plain, spl29_58, inference(avatar_contradiction_clause, [], [f3268])).
fof(f3268, plain, ($false | spl29_58), inference(subsumption_resolution, [], [f3267, f457])).
fof(f3267, plain, (~ l2_lattices(sK1) | spl29_58), inference(resolution, [], [f1186, f325])).
fof(f325, plain, ! [X0] : (v1_funct_1(u2_lattices(X0)) | ~ l2_lattices(X0)), inference(cnf_transformation, [], [f150])).
fof(f1186, plain, (~ v1_funct_1(u2_lattices(sK1)) | spl29_58), inference(avatar_component_clause, [], [f1184])).
fof(f1184, plain, (spl29_58 <=> v1_funct_1(u2_lattices(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_58])])).
fof(f3266, plain, spl29_57, inference(avatar_contradiction_clause, [], [f3265])).
fof(f3265, plain, ($false | spl29_57), inference(subsumption_resolution, [], [f3264, f450])).
fof(f450, plain, l1_lattices(sK1), inference(resolution, [], [f312, f274])).
fof(f312, plain, ! [X0] : (~ l3_lattices(X0) | l1_lattices(X0)), inference(cnf_transformation, [], [f139])).
fof(f3264, plain, (~ l1_lattices(sK1) | spl29_57), inference(resolution, [], [f1182, f711])).
fof(f711, plain, ! [X1] : (m1_relset_1(u1_lattices(X1), k2_zfmisc_1(u1_struct_0(X1), u1_struct_0(X1)), u1_struct_0(X1)) | ~ l1_lattices(X1)), inference(resolution, [], [f324, f425])).
fof(f324, plain, ! [X0] : (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ! [X0] : ((m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0))) | ~ l1_lattices(X0)), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ! [X0] : (l1_lattices(X0) => (m2_relset_1(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) & v1_funct_1(u1_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', dt_u1_lattices)).
fof(f1182, plain, (~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl29_57), inference(avatar_component_clause, [], [f1180])).
fof(f1180, plain, (spl29_57 <=> m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_57])])).
fof(f3256, plain, spl29_56, inference(avatar_contradiction_clause, [], [f3255])).
fof(f3255, plain, ($false | spl29_56), inference(subsumption_resolution, [], [f3254, f450])).
fof(f3254, plain, (~ l1_lattices(sK1) | spl29_56), inference(resolution, [], [f1178, f323])).
fof(f323, plain, ! [X0] : (v1_funct_2(u1_lattices(X0), k2_zfmisc_1(u1_struct_0(X0), u1_struct_0(X0)), u1_struct_0(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f149])).
fof(f1178, plain, (~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | spl29_56), inference(avatar_component_clause, [], [f1176])).
fof(f1176, plain, (spl29_56 <=> v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_56])])).
fof(f2852, plain, spl29_119, inference(avatar_contradiction_clause, [], [f2851])).
fof(f2851, plain, ($false | spl29_119), inference(subsumption_resolution, [], [f2850, f272])).
fof(f2850, plain, (v3_struct_0(sK1) | spl29_119), inference(subsumption_resolution, [], [f2849, f273])).
fof(f2849, plain, (~ v10_lattices(sK1) | v3_struct_0(sK1) | spl29_119), inference(subsumption_resolution, [], [f2848, f274])).
fof(f2848, plain, (~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | spl29_119), inference(subsumption_resolution, [], [f2847, f275])).
fof(f2847, plain, (~ v14_lattices(sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1) | spl29_119), inference(resolution, [], [f2835, f439])).
fof(f2835, plain, (~ v13_lattices(k1_lattice2(sK1)) | spl29_119), inference(avatar_component_clause, [], [f2833])).
fof(f1209, plain, spl29_55, inference(avatar_contradiction_clause, [], [f1208])).
fof(f1208, plain, ($false | spl29_55), inference(subsumption_resolution, [], [f1207, f450])).
fof(f1207, plain, (~ l1_lattices(sK1) | spl29_55), inference(resolution, [], [f1174, f322])).
fof(f322, plain, ! [X0] : (v1_funct_1(u1_lattices(X0)) | ~ l1_lattices(X0)), inference(cnf_transformation, [], [f149])).
fof(f1174, plain, (~ v1_funct_1(u1_lattices(sK1)) | spl29_55), inference(avatar_component_clause, [], [f1172])).
fof(f1172, plain, (spl29_55 <=> v1_funct_1(u1_lattices(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_55])])).
fof(f1206, plain, (~ spl29_55 | ~ spl29_56 | ~ spl29_57 | ~ spl29_58 | ~ spl29_59 | ~ spl29_60 | spl29_63), inference(avatar_split_clause, [], [f1170, f1204, f1192, f1188, f1184, f1180, f1176, f1172])).
fof(f1170, plain, ! [X17, X15, X16] : (~ (k1_lattice2(sK1) = g3_lattices(X15, X16, X17)) | (u1_struct_0(sK1) = X15) | ~ m1_relset_1(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u2_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u2_lattices(sK1)) | ~ m1_relset_1(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_2(u1_lattices(sK1), k2_zfmisc_1(u1_struct_0(sK1), u1_struct_0(sK1)), u1_struct_0(sK1)) | ~ v1_funct_1(u1_lattices(sK1))), inference(superposition, [], [f371, f684])).
fof(f684, plain, (k1_lattice2(sK1) = g3_lattices(u1_struct_0(sK1), u1_lattices(sK1), u2_lattices(sK1))), inference(resolution, [], [f299, f274])).
fof(f299, plain, ! [X0] : (~ l3_lattices(X0) | (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ! [X0] : ((k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0))) | ~ l3_lattices(X0)), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ! [X0] : (l3_lattices(X0) => (k1_lattice2(X0) = g3_lattices(u1_struct_0(X0), u1_lattices(X0), u2_lattices(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', d2_lattice2)).
fof(f371, plain, ! [X4, X2, X0, X5, X3, X1] : (~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) | (X0 = X3) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | ~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1)), inference(flattening, [], [f181])).
fof(f181, plain, ! [X0, X1, X2] : (! [X3, X4, X5] : (((X2 = X5) & (X1 = X4) & (X0 = X3)) | ~ (g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5))) | (~ m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X2) | ~ m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) | ~ v1_funct_1(X1))), inference(ennf_transformation, [], [f62])).
fof(f62, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X2, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X2) & m1_relset_1(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_2(X1, k2_zfmisc_1(X0, X0), X0) & v1_funct_1(X1)) => ! [X3, X4, X5] : ((g3_lattices(X0, X1, X2) = g3_lattices(X3, X4, X5)) => ((X2 = X5) & (X1 = X4) & (X0 = X3)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', free_g3_lattices)).
fof(f883, plain, ~ spl29_11, inference(avatar_contradiction_clause, [], [f882])).
fof(f882, plain, ($false | ~ spl29_11), inference(subsumption_resolution, [], [f878, f599])).
fof(f599, plain, sP0(sK1), inference(subsumption_resolution, [], [f598, f272])).
fof(f598, plain, (sP0(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f592, f273])).
fof(f592, plain, (sP0(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(resolution, [], [f370, f274])).
fof(f370, plain, ! [X0] : (~ l3_lattices(X0) | sP0(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ! [X0] : (sP0(X0) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(definition_folding, [], [f180, e209])).
fof(f209, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(usedef, [], [e209])).
fof(e209, plain, ! [X0] : (sP0(X0) <=> (v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f180, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0)), inference(flattening, [], [f179])).
fof(f179, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | (~ l3_lattices(X0) | ~ v10_lattices(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ! [X0] : ((l3_lattices(X0) & v10_lattices(X0) & ~ v3_struct_0(X0)) => (v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT307+1.p', fc6_lattice2)).
fof(f878, plain, (~ sP0(sK1) | ~ spl29_11), inference(resolution, [], [f811, f361])).
fof(f361, plain, ! [X0] : (~ v3_struct_0(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f236])).
fof(f236, plain, ! [X0] : ((v10_lattices(k1_lattice2(X0)) & v9_lattices(k1_lattice2(X0)) & v8_lattices(k1_lattice2(X0)) & v7_lattices(k1_lattice2(X0)) & v6_lattices(k1_lattice2(X0)) & v5_lattices(k1_lattice2(X0)) & v4_lattices(k1_lattice2(X0)) & v3_lattices(k1_lattice2(X0)) & ~ v3_struct_0(k1_lattice2(X0))) | ~ sP0(X0)), inference(nnf_transformation, [], [f209])).
fof(f811, plain, (v3_struct_0(k1_lattice2(sK1)) | ~ spl29_11), inference(avatar_component_clause, [], [f809])).
fof(f860, plain, spl29_13, inference(avatar_contradiction_clause, [], [f859])).
fof(f859, plain, ($false | spl29_13), inference(subsumption_resolution, [], [f858, f274])).
fof(f858, plain, (~ l3_lattices(sK1) | spl29_13), inference(resolution, [], [f819, f307])).
fof(f819, plain, (~ l3_lattices(k1_lattice2(sK1)) | spl29_13), inference(avatar_component_clause, [], [f817])).
fof(f850, plain, spl29_12, inference(avatar_contradiction_clause, [], [f849])).
fof(f849, plain, ($false | spl29_12), inference(subsumption_resolution, [], [f848, f599])).
fof(f848, plain, (~ sP0(sK1) | spl29_12), inference(resolution, [], [f815, f369])).
fof(f369, plain, ! [X0] : (v10_lattices(k1_lattice2(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f236])).
fof(f815, plain, (~ v10_lattices(k1_lattice2(sK1)) | spl29_12), inference(avatar_component_clause, [], [f813])).
fof(f824, plain, (spl29_11 | ~ spl29_12 | ~ spl29_13 | spl29_14), inference(avatar_split_clause, [], [f805, f821, f817, f813, f809])).
fof(f805, plain, (m1_filter_0(k7_filter_2(sK1, sK2), k1_lattice2(sK1)) | ~ l3_lattices(k1_lattice2(sK1)) | ~ v10_lattices(k1_lattice2(sK1)) | v3_struct_0(k1_lattice2(sK1))), inference(resolution, [], [f791, f423])).
fof(f791, plain, m1_filter_2(k7_filter_2(sK1, sK2), k1_lattice2(sK1)), inference(subsumption_resolution, [], [f790, f272])).
fof(f790, plain, (m1_filter_2(k7_filter_2(sK1, sK2), k1_lattice2(sK1)) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f789, f273])).
fof(f789, plain, (m1_filter_2(k7_filter_2(sK1, sK2), k1_lattice2(sK1)) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f788, f274])).
fof(f788, plain, (m1_filter_2(k7_filter_2(sK1, sK2), k1_lattice2(sK1)) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(subsumption_resolution, [], [f787, f276])).
fof(f787, plain, (m1_filter_2(k7_filter_2(sK1, sK2), k1_lattice2(sK1)) | ~ m2_filter_2(sK2, sK1) | ~ l3_lattices(sK1) | ~ v10_lattices(sK1) | v3_struct_0(sK1)), inference(superposition, [], [f305, f786])).