fof(f12455, plain, $false, inference(avatar_sat_refutation, [], [f1295, f2225, f2781, f2830, f3216, f3333, f6312, f6330, f6348, f6518, f12388, f12413, f12440, f12454])).
fof(f12454, plain, spl41_72, inference(avatar_contradiction_clause, [], [f12453])).
fof(f12453, plain, ($false | spl41_72), inference(subsumption_resolution, [], [f12452, f505])).
fof(f505, plain, v2_orders_2(sK9), inference(cnf_transformation, [], [f431])).
fof(f431, plain, ((~ v1_waybel34(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) & v4_waybel_0(k2_yellow_2(sK9, sK9, sK10), sK9) & m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v7_waybel_1(sK10, sK9) & v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(sK10)) & l1_orders_2(sK9) & v3_lattice3(sK9) & v2_lattice3(sK9) & v1_lattice3(sK9) & v4_orders_2(sK9) & v3_orders_2(sK9) & v2_orders_2(sK9)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10])], [f219, f430, f429])).
fof(f429, plain, (? [X0] : (? [X1] : (~ v1_waybel34(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) & v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (? [X1] : (~ v1_waybel34(k2_waybel_1(sK9, sK9, X1), sK9, k2_yellow_2(sK9, sK9, X1)) & v4_waybel_0(k2_yellow_2(sK9, sK9, X1), sK9) & m2_relset_1(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v7_waybel_1(X1, sK9) & v1_funct_2(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(X1)) & l1_orders_2(sK9) & v3_lattice3(sK9) & v2_lattice3(sK9) & v1_lattice3(sK9) & v4_orders_2(sK9) & v3_orders_2(sK9) & v2_orders_2(sK9))), introduced(choice_axiom, [])).
fof(f430, plain, (? [X1] : (~ v1_waybel34(k2_waybel_1(sK9, sK9, X1), sK9, k2_yellow_2(sK9, sK9, X1)) & v4_waybel_0(k2_yellow_2(sK9, sK9, X1), sK9) & m2_relset_1(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v7_waybel_1(X1, sK9) & v1_funct_2(X1, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(X1)) => (~ v1_waybel34(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) & v4_waybel_0(k2_yellow_2(sK9, sK9, sK10), sK9) & m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v7_waybel_1(sK10, sK9) & v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) & v1_funct_1(sK10))), introduced(choice_axiom, [])).
fof(f219, plain, ? [X0] : (? [X1] : (~ v1_waybel34(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) & v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) & m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)), inference(flattening, [], [f218])).
fof(f218, plain, ? [X0] : (? [X1] : ((~ v1_waybel34(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) & v4_waybel_0(k2_yellow_2(X0, X0, X1), X0)) & (m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1))) & (l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) => v1_waybel34(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) => v1_waybel34(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', t42_waybel34)).
fof(f12452, plain, (~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12451, f506])).
fof(f506, plain, v3_orders_2(sK9), inference(cnf_transformation, [], [f431])).
fof(f12451, plain, (~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12450, f507])).
fof(f507, plain, v4_orders_2(sK9), inference(cnf_transformation, [], [f431])).
fof(f12450, plain, (~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12449, f508])).
fof(f508, plain, v1_lattice3(sK9), inference(cnf_transformation, [], [f431])).
fof(f12449, plain, (~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12448, f509])).
fof(f509, plain, v2_lattice3(sK9), inference(cnf_transformation, [], [f431])).
fof(f12448, plain, (~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12447, f510])).
fof(f510, plain, v3_lattice3(sK9), inference(cnf_transformation, [], [f431])).
fof(f12447, plain, (~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12446, f511])).
fof(f511, plain, l1_orders_2(sK9), inference(cnf_transformation, [], [f431])).
fof(f12446, plain, (~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12445, f512])).
fof(f512, plain, v1_funct_1(sK10), inference(cnf_transformation, [], [f431])).
fof(f12445, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12444, f513])).
fof(f513, plain, v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)), inference(cnf_transformation, [], [f431])).
fof(f12444, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12443, f514])).
fof(f514, plain, v7_waybel_1(sK10, sK9), inference(cnf_transformation, [], [f431])).
fof(f12443, plain, (~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12442, f515])).
fof(f515, plain, m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)), inference(cnf_transformation, [], [f431])).
fof(f12442, plain, (~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(subsumption_resolution, [], [f12441, f516])).
fof(f516, plain, v4_waybel_0(k2_yellow_2(sK9, sK9, sK10), sK9), inference(cnf_transformation, [], [f431])).
fof(f12441, plain, (~ v4_waybel_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_72), inference(resolution, [], [f6517, f1008])).
fof(f1008, plain, ! [X0, X1] : (v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) | ~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f504])).
fof(f504, plain, ! [X0] : (! [X1] : (((v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) | ~ v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)) & (v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) | ~ v4_waybel_0(k2_yellow_2(X0, X0, X1), X0))) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(nnf_transformation, [], [f404])).
fof(f404, plain, ! [X0] : (! [X1] : ((v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f403])).
fof(f403, plain, ! [X0] : (! [X1] : ((v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)) | (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f141])).
fof(f141, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v4_waybel_0(k2_yellow_2(X0, X0, X1), X0) <=> v22_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', t40_waybel34)).
fof(f6517, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | spl41_72), inference(avatar_component_clause, [], [f6515])).
fof(f6515, plain, (spl41_72 <=> v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9)), introduced(avatar_definition, [new_symbols(naming, [spl41_72])])).
fof(f12440, plain, spl41_59, inference(avatar_contradiction_clause, [], [f12439])).
fof(f12439, plain, ($false | spl41_59), inference(subsumption_resolution, [], [f12438, f505])).
fof(f12438, plain, (~ v2_orders_2(sK9) | spl41_59), inference(subsumption_resolution, [], [f12437, f506])).
fof(f12437, plain, (~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_59), inference(subsumption_resolution, [], [f12436, f507])).
fof(f12436, plain, (~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_59), inference(subsumption_resolution, [], [f12435, f508])).
fof(f12435, plain, (~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_59), inference(subsumption_resolution, [], [f12434, f509])).
fof(f12434, plain, (~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_59), inference(subsumption_resolution, [], [f12433, f510])).
fof(f12433, plain, (~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_59), inference(subsumption_resolution, [], [f12432, f511])).
fof(f12432, plain, (~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_59), inference(subsumption_resolution, [], [f12431, f512])).
fof(f12431, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_59), inference(subsumption_resolution, [], [f12430, f513])).
fof(f12430, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_59), inference(subsumption_resolution, [], [f12429, f514])).
fof(f12429, plain, (~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_59), inference(subsumption_resolution, [], [f12428, f515])).
fof(f12428, plain, (~ m2_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_59), inference(resolution, [], [f6243, f1003])).
fof(f1003, plain, ! [X0, X1] : (v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1)) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f401])).
fof(f401, plain, ! [X0] : (! [X1] : (((k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) & v17_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) | ~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f400])).
fof(f400, plain, ! [X0] : (! [X1] : (((k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) & v17_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))) | (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f139])).
fof(f139, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => ((k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) & (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) & v17_waybel_0(k3_waybel_1(X0, X0, X1), k2_yellow_2(X0, X0, X1), X0) & v18_waybel_0(k2_waybel_1(X0, X0, X1), X0, k2_yellow_2(X0, X0, X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', t39_waybel34)).
fof(f6243, plain, (~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | spl41_59), inference(avatar_component_clause, [], [f6241])).
fof(f6241, plain, (spl41_59 <=> v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl41_59])])).
fof(f12413, plain, (spl41_5 | spl41_60), inference(avatar_contradiction_clause, [], [f12412])).
fof(f12412, plain, ($false | (spl41_5 | spl41_60)), inference(subsumption_resolution, [], [f12411, f1342])).
fof(f1342, plain, m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)), inference(resolution, [], [f997, f515])).
fof(f997, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f503])).
fof(f503, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f134])).
fof(f134, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', redefinition_m2_relset_1)).
fof(f12411, plain, (~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | (spl41_5 | spl41_60)), inference(subsumption_resolution, [], [f12410, f1204])).
fof(f1204, plain, (~ v3_struct_0(sK9) | spl41_5), inference(avatar_component_clause, [], [f1203])).
fof(f1203, plain, (spl41_5 <=> v3_struct_0(sK9)), introduced(avatar_definition, [new_symbols(naming, [spl41_5])])).
fof(f12410, plain, (v3_struct_0(sK9) | ~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | spl41_60), inference(subsumption_resolution, [], [f12409, f511])).
fof(f12409, plain, (~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | spl41_60), inference(subsumption_resolution, [], [f12408, f512])).
fof(f12408, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | spl41_60), inference(subsumption_resolution, [], [f12407, f513])).
fof(f12407, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | spl41_60), inference(duplicate_literal_removal, [], [f12406])).
fof(f12406, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | spl41_60), inference(resolution, [], [f6247, f3283])).
fof(f3283, plain, ! [X10, X11, X9] : (m1_relset_1(k2_waybel_1(X10, X11, X9), u1_struct_0(X10), u1_struct_0(k2_yellow_2(X10, X11, X9))) | ~ v1_funct_2(X9, u1_struct_0(X10), u1_struct_0(X11)) | ~ v1_funct_1(X9) | ~ l1_orders_2(X11) | v3_struct_0(X11) | ~ l1_orders_2(X10) | v3_struct_0(X10) | ~ m1_relset_1(X9, u1_struct_0(X10), u1_struct_0(X11))), inference(resolution, [], [f668, f997])).
fof(f668, plain, ! [X2, X0, X1] : (m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f307])).
fof(f307, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f306])).
fof(f306, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f62])).
fof(f62, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m2_relset_1(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) & v1_funct_1(k2_waybel_1(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', dt_k2_waybel_1)).
fof(f6247, plain, (~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | spl41_60), inference(avatar_component_clause, [], [f6245])).
fof(f6245, plain, (spl41_60 <=> m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))), introduced(avatar_definition, [new_symbols(naming, [spl41_60])])).
fof(f12388, plain, (spl41_5 | spl41_58), inference(avatar_contradiction_clause, [], [f12387])).
fof(f12387, plain, ($false | (spl41_5 | spl41_58)), inference(subsumption_resolution, [], [f12386, f1204])).
fof(f12386, plain, (v3_struct_0(sK9) | spl41_58), inference(subsumption_resolution, [], [f12385, f511])).
fof(f12385, plain, (~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl41_58), inference(subsumption_resolution, [], [f12384, f512])).
fof(f12384, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl41_58), inference(subsumption_resolution, [], [f12383, f513])).
fof(f12383, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl41_58), inference(subsumption_resolution, [], [f12382, f1342])).
fof(f12382, plain, (~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl41_58), inference(duplicate_literal_removal, [], [f12381])).
fof(f12381, plain, (~ m1_relset_1(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | spl41_58), inference(resolution, [], [f6239, f667])).
fof(f667, plain, ! [X2, X0, X1] : (v1_funct_2(k2_waybel_1(X0, X1, X2), u1_struct_0(X0), u1_struct_0(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f307])).
fof(f6239, plain, (~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | spl41_58), inference(avatar_component_clause, [], [f6237])).
fof(f6237, plain, (spl41_58 <=> v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)))), introduced(avatar_definition, [new_symbols(naming, [spl41_58])])).
fof(f6518, plain, (~ spl41_66 | ~ spl41_64 | ~ spl41_65 | ~ spl41_72 | spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_38 | ~ spl41_39), inference(avatar_split_clause, [], [f6513, f2833, f2827, f2814, f2364, f1203, f6515, f6327, f6309, f6345])).
fof(f6345, plain, (spl41_66 <=> v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9))), introduced(avatar_definition, [new_symbols(naming, [spl41_66])])).
fof(f6309, plain, (spl41_64 <=> v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9)), introduced(avatar_definition, [new_symbols(naming, [spl41_64])])).
fof(f6327, plain, (spl41_65 <=> m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9))), introduced(avatar_definition, [new_symbols(naming, [spl41_65])])).
fof(f2364, plain, (spl41_29 <=> l1_orders_2(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl41_29])])).
fof(f2814, plain, (spl41_36 <=> v3_lattice3(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl41_36])])).
fof(f2827, plain, (spl41_38 <=> v1_lattice3(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl41_38])])).
fof(f2833, plain, (spl41_39 <=> v2_lattice3(k2_yellow_2(sK9, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl41_39])])).
fof(f6513, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | (spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6512, f2671])).
fof(f2671, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f2670, f505])).
fof(f2670, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f2669, f511])).
fof(f2669, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f2662, f2645])).
fof(f2645, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl41_5), inference(subsumption_resolution, [], [f2644, f1204])).
fof(f2644, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2643, f511])).
fof(f2643, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2642, f512])).
fof(f2642, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2641, f513])).
fof(f2641, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2624])).
fof(f2624, plain, (m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f671, f1342])).
fof(f671, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f309])).
fof(f309, plain, ! [X0, X1, X2] : ((m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f308])).
fof(f308, plain, ! [X0, X1, X2] : ((m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f63])).
fof(f63, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m1_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', dt_k2_yellow_2)).
fof(f2662, plain, (v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_5), inference(resolution, [], [f2609, f639])).
fof(f639, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v2_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f284])).
fof(f284, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v2_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f283])).
fof(f283, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v2_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f48])).
fof(f48, plain, ! [X0] : ((l1_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v2_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', cc6_yellow_0)).
fof(f2609, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl41_5), inference(subsumption_resolution, [], [f2608, f1204])).
fof(f2608, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2607, f511])).
fof(f2607, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2606, f512])).
fof(f2606, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2605, f513])).
fof(f2605, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2588])).
fof(f2588, plain, (v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f670, f1342])).
fof(f670, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f309])).
fof(f6512, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6511, f2668])).
fof(f2668, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f2667, f506])).
fof(f2667, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f2666, f511])).
fof(f2666, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f2661, f2645])).
fof(f2661, plain, (v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v3_orders_2(sK9) | spl41_5), inference(resolution, [], [f2609, f644])).
fof(f644, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v3_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v3_orders_2(X0)), inference(cnf_transformation, [], [f288])).
fof(f288, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v3_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v3_orders_2(X0)), inference(flattening, [], [f287])).
fof(f287, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v3_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v3_orders_2(X0))), inference(ennf_transformation, [], [f50])).
fof(f50, plain, ! [X0] : ((l1_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v3_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', cc7_yellow_0)).
fof(f6511, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6510, f2665])).
fof(f2665, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f2664, f507])).
fof(f2664, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f2663, f511])).
fof(f2663, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f2660, f2645])).
fof(f2660, plain, (v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | spl41_5), inference(resolution, [], [f2609, f649])).
fof(f649, plain, ! [X0, X1] : (~ v4_yellow_0(X1, X0) | v4_orders_2(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0)), inference(cnf_transformation, [], [f292])).
fof(f292, plain, ! [X0] : (! [X1] : ((v4_yellow_0(X1, X0) & v4_orders_2(X1)) | ~ v4_yellow_0(X1, X0) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0)), inference(flattening, [], [f291])).
fof(f291, plain, ! [X0] : (! [X1] : (((v4_yellow_0(X1, X0) & v4_orders_2(X1)) | ~ v4_yellow_0(X1, X0)) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v4_orders_2(X0))), inference(ennf_transformation, [], [f52])).
fof(f52, plain, ! [X0] : ((l1_orders_2(X0) & v4_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) => (v4_yellow_0(X1, X0) & v4_orders_2(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', cc8_yellow_0)).
fof(f6510, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6509, f2829])).
fof(f2829, plain, (v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl41_38), inference(avatar_component_clause, [], [f2827])).
fof(f6509, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_39)), inference(subsumption_resolution, [], [f6508, f2835])).
fof(f2835, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl41_39), inference(avatar_component_clause, [], [f2833])).
fof(f6508, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_36)), inference(subsumption_resolution, [], [f6507, f2815])).
fof(f2815, plain, (v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl41_36), inference(avatar_component_clause, [], [f2814])).
fof(f6507, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29)), inference(subsumption_resolution, [], [f6506, f2365])).
fof(f2365, plain, (l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ spl41_29), inference(avatar_component_clause, [], [f2364])).
fof(f6506, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6505, f505])).
fof(f6505, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6504, f506])).
fof(f6504, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6503, f507])).
fof(f6503, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6502, f508])).
fof(f6502, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6501, f509])).
fof(f6501, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6500, f510])).
fof(f6500, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6499, f511])).
fof(f6499, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6498, f2393])).
fof(f2393, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f2392, f1204])).
fof(f2392, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2391, f511])).
fof(f2391, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2390, f512])).
fof(f2390, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2389, f513])).
fof(f2389, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2372])).
fof(f2372, plain, (v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f672, f1342])).
fof(f672, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v1_funct_1(k3_waybel_1(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f311])).
fof(f311, plain, ! [X0, X1, X2] : ((m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f310])).
fof(f310, plain, ! [X0, X1, X2] : ((m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f65])).
fof(f65, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (m2_relset_1(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_2(k3_waybel_1(X0, X1, X2), u1_struct_0(k2_yellow_2(X0, X1, X2)), u1_struct_0(X1)) & v1_funct_1(k3_waybel_1(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', dt_k3_waybel_1)).
fof(f6498, plain, (~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), inference(subsumption_resolution, [], [f6384, f517])).
fof(f517, plain, ~ v1_waybel34(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)), inference(cnf_transformation, [], [f431])).
fof(f6384, plain, (v1_waybel34(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v22_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ v1_funct_1(k3_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), inference(superposition, [], [f1001, f3884])).
fof(f3884, plain, (k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))), inference(subsumption_resolution, [], [f3883, f505])).
fof(f3883, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3882, f506])).
fof(f3882, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3881, f507])).
fof(f3881, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3880, f508])).
fof(f3880, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3879, f509])).
fof(f3879, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3878, f510])).
fof(f3878, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3877, f511])).
fof(f3877, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3876, f512])).
fof(f3876, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3875, f513])).
fof(f3875, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3868, f514])).
fof(f3868, plain, ((k2_waybel_1(sK9, sK9, sK10) = k1_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k3_waybel_1(sK9, sK9, sK10))) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(resolution, [], [f1006, f515])).
fof(f1006, plain, ! [X0, X1] : (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | (k2_waybel_1(X0, X0, X1) = k1_waybel34(k2_yellow_2(X0, X0, X1), X0, k3_waybel_1(X0, X0, X1))) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f401])).
fof(f1001, plain, ! [X2, X0, X1] : (v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ v22_waybel_0(X2, X0, X1) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f397])).
fof(f397, plain, ! [X0] : (! [X1] : (! [X2] : (v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ v22_waybel_0(X2, X0, X1) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1)) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f396])).
fof(f396, plain, ! [X0] : (! [X1] : (! [X2] : ((v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0) | ~ v22_waybel_0(X2, X0, X1)) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v17_waybel_0(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1))) | (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f137])).
fof(f137, plain, ! [X0] : ((l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => ! [X1] : ((l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v17_waybel_0(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => (v22_waybel_0(X2, X0, X1) => v1_waybel34(k1_waybel34(X0, X1, X2), X1, X0))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', t24_waybel34)).
fof(f6348, plain, (~ spl41_58 | ~ spl41_60 | spl41_66 | spl41_5 | ~ spl41_29 | ~ spl41_38 | ~ spl41_39), inference(avatar_split_clause, [], [f6343, f2833, f2827, f2364, f1203, f6345, f6245, f6237])).
fof(f6343, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | (spl41_5 | ~ spl41_29 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6342, f2671])).
fof(f6342, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6341, f2668])).
fof(f6341, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6340, f2665])).
fof(f6340, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6339, f2829])).
fof(f6339, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_39)), inference(subsumption_resolution, [], [f6338, f2835])).
fof(f6338, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29)), inference(subsumption_resolution, [], [f6337, f2365])).
fof(f6337, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6336, f505])).
fof(f6336, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6335, f506])).
fof(f6335, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6334, f507])).
fof(f6334, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6333, f508])).
fof(f6333, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6332, f509])).
fof(f6332, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6331, f511])).
fof(f6331, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6220, f2311])).
fof(f2311, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f2310, f1204])).
fof(f2310, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2309, f511])).
fof(f2309, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2308, f512])).
fof(f2308, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2307, f513])).
fof(f2307, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2290])).
fof(f2290, plain, (v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f666, f1342])).
fof(f666, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | v1_funct_1(k2_waybel_1(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f307])).
fof(f6220, plain, (v1_funct_2(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), inference(superposition, [], [f664, f3863])).
fof(f3863, plain, (k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))), inference(subsumption_resolution, [], [f3862, f505])).
fof(f3862, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3861, f506])).
fof(f3861, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3860, f507])).
fof(f3860, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3859, f508])).
fof(f3859, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3858, f509])).
fof(f3858, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3857, f510])).
fof(f3857, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3856, f511])).
fof(f3856, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3855, f512])).
fof(f3855, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3854, f513])).
fof(f3854, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(subsumption_resolution, [], [f3847, f514])).
fof(f3847, plain, ((k3_waybel_1(sK9, sK9, sK10) = k2_waybel34(k2_yellow_2(sK9, sK9, sK10), sK9, k2_waybel_1(sK9, sK9, sK10))) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(resolution, [], [f1005, f515])).
fof(f1005, plain, ! [X0, X1] : (~ m2_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | (k3_waybel_1(X0, X0, X1) = k2_waybel34(k2_yellow_2(X0, X0, X1), X0, k2_waybel_1(X0, X0, X1))) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f401])).
fof(f664, plain, ! [X2, X0, X1] : (v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f305])).
fof(f305, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f304])).
fof(f304, plain, ! [X0, X1, X2] : ((m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k2_waybel34(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', dt_k2_waybel34)).
fof(f6330, plain, (~ spl41_58 | ~ spl41_60 | spl41_65 | spl41_5 | ~ spl41_29 | ~ spl41_38 | ~ spl41_39), inference(avatar_split_clause, [], [f6325, f2833, f2827, f2364, f1203, f6327, f6245, f6237])).
fof(f6325, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | (spl41_5 | ~ spl41_29 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6324, f2671])).
fof(f6324, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6323, f2668])).
fof(f6323, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6322, f2665])).
fof(f6322, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6321, f2829])).
fof(f6321, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_39)), inference(subsumption_resolution, [], [f6320, f2835])).
fof(f6320, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29)), inference(subsumption_resolution, [], [f6319, f2365])).
fof(f6319, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6318, f505])).
fof(f6318, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6317, f506])).
fof(f6317, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6316, f507])).
fof(f6316, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6315, f508])).
fof(f6315, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6314, f509])).
fof(f6314, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6313, f511])).
fof(f6313, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6219, f2311])).
fof(f6219, plain, (m2_relset_1(k3_waybel_1(sK9, sK9, sK10), u1_struct_0(k2_yellow_2(sK9, sK9, sK10)), u1_struct_0(sK9)) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), inference(superposition, [], [f665, f3863])).
fof(f665, plain, ! [X2, X0, X1] : (m2_relset_1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f305])).
fof(f6312, plain, (~ spl41_58 | ~ spl41_59 | ~ spl41_60 | spl41_64 | spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_38 | ~ spl41_39), inference(avatar_split_clause, [], [f6307, f2833, f2827, f2814, f2364, f1203, f6309, f6245, f6241, f6237])).
fof(f6307, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | (spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6306, f2671])).
fof(f6306, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6305, f2668])).
fof(f6305, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6304, f2665])).
fof(f6304, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_38 | ~ spl41_39)), inference(subsumption_resolution, [], [f6303, f2829])).
fof(f6303, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_36 | ~ spl41_39)), inference(subsumption_resolution, [], [f6302, f2835])).
fof(f6302, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29 | ~ spl41_36)), inference(subsumption_resolution, [], [f6301, f2815])).
fof(f6301, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29)), inference(subsumption_resolution, [], [f6300, f2365])).
fof(f6300, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6299, f505])).
fof(f6299, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6298, f506])).
fof(f6298, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6297, f507])).
fof(f6297, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6296, f508])).
fof(f6296, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6295, f509])).
fof(f6295, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6294, f510])).
fof(f6294, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6293, f511])).
fof(f6293, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f6218, f2311])).
fof(f6218, plain, (v17_waybel_0(k3_waybel_1(sK9, sK9, sK10), k2_yellow_2(sK9, sK9, sK10), sK9) | ~ m1_relset_1(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v18_waybel_0(k2_waybel_1(sK9, sK9, sK10), sK9, k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(k2_waybel_1(sK9, sK9, sK10), u1_struct_0(sK9), u1_struct_0(k2_yellow_2(sK9, sK9, sK10))) | ~ v1_funct_1(k2_waybel_1(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | ~ l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_orders_2(k2_yellow_2(sK9, sK9, sK10))), inference(superposition, [], [f746, f3863])).
fof(f746, plain, ! [X2, X0, X1] : (v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f336])).
fof(f336, plain, ! [X0, X1, X2] : ((v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f335])).
fof(f335, plain, ! [X0, X1, X2] : ((v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v18_waybel_0(X2, X1, X0) | ~ v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | ~ v3_lattice3(X1) | ~ v2_lattice3(X1) | ~ v1_lattice3(X1) | ~ v4_orders_2(X1) | ~ v3_orders_2(X1) | ~ v2_orders_2(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f185])).
fof(f185, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f169])).
fof(f169, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_partfun1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2)))), inference(pure_predicate_removal, [], [f90])).
fof(f90, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X1), u1_struct_0(X0)) & v18_waybel_0(X2, X1, X0) & v1_funct_2(X2, u1_struct_0(X1), u1_struct_0(X0)) & v1_funct_1(X2) & l1_orders_2(X1) & v3_lattice3(X1) & v2_lattice3(X1) & v1_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & v2_orders_2(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v5_orders_3(k2_waybel34(X0, X1, X2), X0, X1) & v4_waybel_1(k2_waybel34(X0, X1, X2), X0, X1) & v21_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v19_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v17_waybel_0(k2_waybel34(X0, X1, X2), X0, X1) & v1_partfun1(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(k2_waybel34(X0, X1, X2), u1_struct_0(X0), u1_struct_0(X1)) & ~ v1_xboole_0(k2_waybel34(X0, X1, X2)) & v1_funct_1(k2_waybel34(X0, X1, X2)) & v1_relat_1(k2_waybel34(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', fc2_waybel34)).
fof(f3333, plain, (~ spl41_19 | spl41_36), inference(avatar_contradiction_clause, [], [f3332])).
fof(f3332, plain, ($false | (~ spl41_19 | spl41_36)), inference(subsumption_resolution, [], [f3331, f505])).
fof(f3331, plain, (~ v2_orders_2(sK9) | (~ spl41_19 | spl41_36)), inference(subsumption_resolution, [], [f3330, f506])).
fof(f3330, plain, (~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl41_19 | spl41_36)), inference(subsumption_resolution, [], [f3329, f507])).
fof(f3329, plain, (~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl41_19 | spl41_36)), inference(subsumption_resolution, [], [f3328, f508])).
fof(f3328, plain, (~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl41_19 | spl41_36)), inference(subsumption_resolution, [], [f3327, f509])).
fof(f3327, plain, (~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl41_19 | spl41_36)), inference(subsumption_resolution, [], [f3326, f510])).
fof(f3326, plain, (~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl41_19 | spl41_36)), inference(subsumption_resolution, [], [f3325, f511])).
fof(f3325, plain, (~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl41_19 | spl41_36)), inference(subsumption_resolution, [], [f3324, f512])).
fof(f3324, plain, (~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl41_19 | spl41_36)), inference(subsumption_resolution, [], [f3323, f513])).
fof(f3323, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | (~ spl41_19 | spl41_36)), inference(subsumption_resolution, [], [f3322, f2045])).
fof(f2045, plain, (v6_waybel_1(sK10, sK9) | ~ spl41_19), inference(avatar_component_clause, [], [f2044])).
fof(f2044, plain, (spl41_19 <=> v6_waybel_1(sK10, sK9)), introduced(avatar_definition, [new_symbols(naming, [spl41_19])])).
fof(f3322, plain, (~ v6_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | spl41_36), inference(subsumption_resolution, [], [f3304, f2901])).
fof(f2901, plain, (~ sP2(sK10, sK9) | spl41_36), inference(resolution, [], [f2816, f710])).
fof(f710, plain, ! [X0, X1] : (v3_lattice3(k2_yellow_2(X1, X1, X0)) | ~ sP2(X0, X1)), inference(cnf_transformation, [], [f447])).
fof(f447, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X1, X1, X0)) & v2_lattice3(k2_yellow_2(X1, X1, X0)) & v1_lattice3(k2_yellow_2(X1, X1, X0)) & v25_waybel_0(k2_yellow_2(X1, X1, X0)) & v24_waybel_0(k2_yellow_2(X1, X1, X0)) & v4_yellow_0(k2_yellow_2(X1, X1, X0), X1) & v3_yellow_0(k2_yellow_2(X1, X1, X0)) & v2_yellow_0(k2_yellow_2(X1, X1, X0)) & v1_yellow_0(k2_yellow_2(X1, X1, X0)) & v4_orders_2(k2_yellow_2(X1, X1, X0)) & v3_orders_2(k2_yellow_2(X1, X1, X0)) & v2_orders_2(k2_yellow_2(X1, X1, X0)) & v1_orders_2(k2_yellow_2(X1, X1, X0)) & ~ v3_struct_0(k2_yellow_2(X1, X1, X0))) | ~ sP2(X0, X1)), inference(rectify, [], [f446])).
fof(f446, plain, ! [X1, X0] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ sP2(X1, X0)), inference(nnf_transformation, [], [f415])).
fof(f415, plain, ! [X1, X0] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ sP2(X1, X0)), inference(usedef, [], [e415])).
fof(e415, plain, ! [X1, X0] : (sP2(X1, X0) <=> (v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f2816, plain, (~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | spl41_36), inference(avatar_component_clause, [], [f2814])).
fof(f3304, plain, (sP2(sK10, sK9) | ~ v6_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v3_lattice3(sK9) | ~ v2_lattice3(sK9) | ~ v1_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9)), inference(resolution, [], [f711, f1342])).
fof(f711, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | sP2(X1, X0) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(cnf_transformation, [], [f416])).
fof(f416, plain, ! [X0, X1] : (sP2(X1, X0) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(definition_folding, [], [f322, e415])).
fof(f322, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0)), inference(flattening, [], [f321])).
fof(f321, plain, ! [X0, X1] : ((v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v6_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v3_lattice3(X0) | ~ v2_lattice3(X0) | ~ v1_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ! [X0, X1] : ((m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1) & l1_orders_2(X0) & v3_lattice3(X0) & v2_lattice3(X0) & v1_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0)) => (v3_lattice3(k2_yellow_2(X0, X0, X1)) & v2_lattice3(k2_yellow_2(X0, X0, X1)) & v1_lattice3(k2_yellow_2(X0, X0, X1)) & v25_waybel_0(k2_yellow_2(X0, X0, X1)) & v24_waybel_0(k2_yellow_2(X0, X0, X1)) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v3_yellow_0(k2_yellow_2(X0, X0, X1)) & v2_yellow_0(k2_yellow_2(X0, X0, X1)) & v1_yellow_0(k2_yellow_2(X0, X0, X1)) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', fc10_waybel34)).
fof(f3216, plain, (spl41_39 | spl41_5), inference(avatar_split_clause, [], [f3215, f1203, f2833])).
fof(f3215, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f3214, f506])).
fof(f3214, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v3_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f3213, f507])).
fof(f3213, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f3212, f509])).
fof(f3212, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v2_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f3211, f511])).
fof(f3211, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f3210, f2645])).
fof(f3210, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f3209, f2429])).
fof(f2429, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f2428, f1204])).
fof(f2428, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2427, f511])).
fof(f2427, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2426, f512])).
fof(f2426, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2425, f513])).
fof(f2425, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(duplicate_literal_removal, [], [f2408])).
fof(f2408, plain, (~ v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | v3_struct_0(sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f754, f1342])).
fof(f754, plain, ! [X2, X0, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v3_struct_0(k2_yellow_2(X0, X1, X2)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f339])).
fof(f339, plain, ! [X0, X1, X2] : ((v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f338])).
fof(f338, plain, ! [X0, X1, X2] : ((v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_orders_2(X1) | v3_struct_0(X1) | ~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f92])).
fof(f92, plain, ! [X0, X1, X2] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_orders_2(X1) & ~ v3_struct_0(X1) & l1_orders_2(X0) & ~ v3_struct_0(X0)) => (v4_yellow_0(k2_yellow_2(X0, X1, X2), X1) & v1_orders_2(k2_yellow_2(X0, X1, X2)) & ~ v3_struct_0(k2_yellow_2(X0, X1, X2)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', fc2_yellow_2)).
fof(f3209, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl41_5), inference(subsumption_resolution, [], [f3206, f2609])).
fof(f3206, plain, (v2_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ v4_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | ~ m1_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v2_lattice3(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | spl41_5), inference(resolution, [], [f3185, f530])).
fof(f530, plain, ! [X0, X1] : (~ v5_yellow_0(X1, X0) | v2_lattice3(X1) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1) | ~ m1_yellow_0(X1, X0) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0)), inference(cnf_transformation, [], [f228])).
fof(f228, plain, ! [X0] : (! [X1] : ((v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v2_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1)) | ~ v5_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0)), inference(flattening, [], [f227])).
fof(f227, plain, ! [X0] : (! [X1] : (((v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v2_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1)) | (~ v5_yellow_0(X1, X0) | ~ v4_yellow_0(X1, X0) | v3_struct_0(X1))) | ~ m1_yellow_0(X1, X0)) | (~ l1_orders_2(X0) | ~ v2_lattice3(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ! [X0] : ((l1_orders_2(X0) & v2_lattice3(X0) & v4_orders_2(X0) & v3_orders_2(X0)) => ! [X1] : (m1_yellow_0(X1, X0) => ((v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & ~ v3_struct_0(X1)) => (v5_yellow_0(X1, X0) & v4_yellow_0(X1, X0) & v2_lattice3(X1) & v4_orders_2(X1) & v3_orders_2(X1) & ~ v3_struct_0(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', cc11_yellow_0)).
fof(f3185, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | spl41_5), inference(subsumption_resolution, [], [f3184, f1204])).
fof(f3184, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3183, f505])).
fof(f3183, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3182, f506])).
fof(f3182, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3181, f507])).
fof(f3181, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3180, f511])).
fof(f3180, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3179, f512])).
fof(f3179, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3178, f513])).
fof(f3178, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f3160, f514])).
fof(f3160, plain, (v5_yellow_0(k2_yellow_2(sK9, sK9, sK10), sK9) | ~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | ~ l1_orders_2(sK9) | ~ v4_orders_2(sK9) | ~ v3_orders_2(sK9) | ~ v2_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f695, f1342])).
fof(f695, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f320])).
fof(f320, plain, ! [X0, X1] : ((v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f319])).
fof(f319, plain, ! [X0, X1] : ((v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1))) | (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ l1_orders_2(X0) | ~ v4_orders_2(X0) | ~ v3_orders_2(X0) | ~ v2_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f177])).
fof(f177, plain, ! [X0, X1] : ((m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1) & l1_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), inference(pure_predicate_removal, [], [f82])).
fof(f82, plain, ! [X0, X1] : ((m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1) & l1_orders_2(X0) & v4_orders_2(X0) & v3_orders_2(X0) & v2_orders_2(X0) & ~ v3_struct_0(X0)) => (v3_waybel_0(k2_yellow_2(X0, X0, X1), X0) & v7_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v5_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_yellow_0(k2_yellow_2(X0, X0, X1), X0) & v4_orders_2(k2_yellow_2(X0, X0, X1)) & v3_orders_2(k2_yellow_2(X0, X0, X1)) & v2_orders_2(k2_yellow_2(X0, X0, X1)) & v1_orders_2(k2_yellow_2(X0, X0, X1)) & ~ v3_struct_0(k2_yellow_2(X0, X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', fc10_waybel10)).
fof(f2830, plain, (spl41_38 | ~ spl41_36 | spl41_5 | ~ spl41_29), inference(avatar_split_clause, [], [f2825, f2364, f1203, f2814, f2827])).
fof(f2825, plain, (~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | (spl41_5 | ~ spl41_29)), inference(subsumption_resolution, [], [f2799, f2429])).
fof(f2799, plain, (~ v3_lattice3(k2_yellow_2(sK9, sK9, sK10)) | v3_struct_0(k2_yellow_2(sK9, sK9, sK10)) | v1_lattice3(k2_yellow_2(sK9, sK9, sK10)) | ~ spl41_29), inference(resolution, [], [f2365, f590])).
fof(f590, plain, ! [X0] : (~ l1_orders_2(X0) | ~ v3_lattice3(X0) | v3_struct_0(X0) | v1_lattice3(X0)), inference(cnf_transformation, [], [f255])).
fof(f255, plain, ! [X0] : ((v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)) | ~ v3_lattice3(X0) | v3_struct_0(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f254])).
fof(f254, plain, ! [X0] : (((v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)) | (~ v3_lattice3(X0) | v3_struct_0(X0))) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ! [X0] : (l1_orders_2(X0) => ((v3_lattice3(X0) & ~ v3_struct_0(X0)) => (v2_lattice3(X0) & v1_lattice3(X0) & ~ v3_struct_0(X0)))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', cc1_yellow_0)).
fof(f2781, plain, (spl41_29 | spl41_5), inference(avatar_split_clause, [], [f2780, f1203, f2364])).
fof(f2780, plain, (l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | spl41_5), inference(subsumption_resolution, [], [f2776, f511])).
fof(f2776, plain, (l1_orders_2(k2_yellow_2(sK9, sK9, sK10)) | ~ l1_orders_2(sK9) | spl41_5), inference(resolution, [], [f2645, f680])).
fof(f680, plain, ! [X0, X1] : (~ m1_yellow_0(X1, X0) | l1_orders_2(X1) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f315])).
fof(f315, plain, ! [X0] : (! [X1] : (l1_orders_2(X1) | ~ m1_yellow_0(X1, X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f72])).
fof(f72, plain, ! [X0] : (l1_orders_2(X0) => ! [X1] : (m1_yellow_0(X1, X0) => l1_orders_2(X1))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', dt_m1_yellow_0)).
fof(f2225, plain, (spl41_19 | spl41_5), inference(avatar_split_clause, [], [f2224, f1203, f2044])).
fof(f2224, plain, (v6_waybel_1(sK10, sK9) | spl41_5), inference(subsumption_resolution, [], [f2223, f1204])).
fof(f2223, plain, (v6_waybel_1(sK10, sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2222, f511])).
fof(f2222, plain, (v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2221, f512])).
fof(f2221, plain, (~ v1_funct_1(sK10) | v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2220, f513])).
fof(f2220, plain, (~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(subsumption_resolution, [], [f2204, f514])).
fof(f2204, plain, (~ v7_waybel_1(sK10, sK9) | ~ v1_funct_2(sK10, u1_struct_0(sK9), u1_struct_0(sK9)) | ~ v1_funct_1(sK10) | v6_waybel_1(sK10, sK9) | ~ l1_orders_2(sK9) | v3_struct_0(sK9)), inference(resolution, [], [f625, f1342])).
fof(f625, plain, ! [X0, X1] : (~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | v6_waybel_1(X1, X0) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ! [X0] : (! [X1] : ((v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1)) | ~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0))) | ~ l1_orders_2(X0) | v3_struct_0(X0)), inference(flattening, [], [f270])).
fof(f270, plain, ! [X0] : (! [X1] : (((v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1)) | (~ v7_waybel_1(X1, X0) | ~ v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) | ~ v1_funct_1(X1))) | ~ m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0))) | (~ l1_orders_2(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f184])).
fof(f184, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) => ((v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v6_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1))))), inference(pure_predicate_removal, [], [f41])).
fof(f41, plain, ! [X0] : ((l1_orders_2(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_relset_1(X1, u1_struct_0(X0), u1_struct_0(X0)) => ((v7_waybel_1(X1, X0) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_1(X1)) => (v6_waybel_1(X1, X0) & v1_partfun1(X1, u1_struct_0(X0), u1_struct_0(X0)) & v1_funct_2(X1, u1_struct_0(X0), u1_struct_0(X0)) & ~ v1_xboole_0(X1) & v1_funct_1(X1))))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', cc3_waybel_1)).
fof(f1295, plain, ~ spl41_5, inference(avatar_contradiction_clause, [], [f1294])).
fof(f1294, plain, ($false | ~ spl41_5), inference(subsumption_resolution, [], [f1293, f511])).
fof(f1293, plain, (~ l1_orders_2(sK9) | ~ spl41_5), inference(subsumption_resolution, [], [f1289, f508])).
fof(f1289, plain, (~ v1_lattice3(sK9) | ~ l1_orders_2(sK9) | ~ spl41_5), inference(resolution, [], [f1205, f570])).
fof(f570, plain, ! [X0] : (~ v3_struct_0(X0) | ~ v1_lattice3(X0) | ~ l1_orders_2(X0)), inference(cnf_transformation, [], [f243])).
fof(f243, plain, ! [X0] : (~ v3_struct_0(X0) | ~ v1_lattice3(X0) | ~ l1_orders_2(X0)), inference(flattening, [], [f242])).
fof(f242, plain, ! [X0] : ((~ v3_struct_0(X0) | ~ v1_lattice3(X0)) | ~ l1_orders_2(X0)), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ! [X0] : (l1_orders_2(X0) => (v1_lattice3(X0) => ~ v3_struct_0(X0))), file('/home/ubuntu/library/tptp/Problems/LAT/LAT369+1.p', cc1_lattice3)).
fof(f1205, plain, (v3_struct_0(sK9) | ~ spl41_5), inference(avatar_component_clause, [], [f1203])).