fof(f106174, plain, $false, inference(avatar_sat_refutation, [], [f675, f679, f687, f692, f697, f705, f709, f716, f720, f724, f728, f733, f738, f1007, f26809, f29421, f29746, f46132, f52270, f52333, f52509, f52765, f54705, f54717, f63417, f63531, f71025, f71092, f71463, f86926, f90482, f90486, f90487, f90841, f91174, f91186, f91194, f95574, f95578, f97071, f97072, f97089, f97391, f97418, f97708, f97742, f97745, f97839, f98165, f100885, f100893, f100894, f104268, f106125])).
fof(f106125, plain, (~ spl96_63 | ~ spl96_392 | spl96_1009 | ~ spl96_1117 | ~ spl96_1118), inference(avatar_contradiction_clause, [], [f106124])).
fof(f106124, plain, ($false | (~ spl96_63 | ~ spl96_392 | spl96_1009 | ~ spl96_1117 | ~ spl96_1118)), inference(subsumption_resolution, [], [f106095, f53817])).
fof(f53817, plain, (~ p2(sK66(sK85)) | spl96_1009), inference(avatar_component_clause, [], [f53816])).
fof(f53816, plain, (spl96_1009 <=> p2(sK66(sK85))), introduced(avatar_definition, [new_symbols(naming, [spl96_1009])])).
fof(f106095, plain, (p2(sK66(sK85)) | (~ spl96_63 | ~ spl96_392 | ~ spl96_1117 | ~ spl96_1118)), inference(resolution, [], [f101286, f54344])).
fof(f54344, plain, (r1(sK65(sK85), sK66(sK85)) | ~ spl96_1117), inference(avatar_component_clause, [], [f54342])).
fof(f54342, plain, (spl96_1117 <=> r1(sK65(sK85), sK66(sK85))), introduced(avatar_definition, [new_symbols(naming, [spl96_1117])])).
fof(f101286, plain, (! [X0] : (~ r1(sK65(sK85), X0) | p2(X0)) | (~ spl96_63 | ~ spl96_392 | ~ spl96_1118)), inference(subsumption_resolution, [], [f101124, f29776])).
fof(f29776, plain, (p2(sK65(sK85)) | ~ spl96_392), inference(avatar_component_clause, [], [f29775])).
fof(f29775, plain, (spl96_392 <=> p2(sK65(sK85))), introduced(avatar_definition, [new_symbols(naming, [spl96_392])])).
fof(f101124, plain, (! [X0] : (~ p2(sK65(sK85)) | ~ r1(sK65(sK85), X0) | p2(X0)) | (~ spl96_63 | ~ spl96_1118)), inference(resolution, [], [f54350, f686])).
fof(f686, plain, (! [X45, X46] : (~ r1(sK85, X45) | ~ p2(X45) | ~ r1(X45, X46) | p2(X46)) | ~ spl96_63), inference(avatar_component_clause, [], [f685])).
fof(f685, plain, (spl96_63 <=> ! [X46, X45] : (~ p2(X45) | ~ r1(sK85, X45) | ~ r1(X45, X46) | p2(X46))), introduced(avatar_definition, [new_symbols(naming, [spl96_63])])).
fof(f54350, plain, (r1(sK85, sK65(sK85)) | ~ spl96_1118), inference(avatar_component_clause, [], [f54348])).
fof(f54348, plain, (spl96_1118 <=> r1(sK85, sK65(sK85))), introduced(avatar_definition, [new_symbols(naming, [spl96_1118])])).
fof(f104268, plain, (spl96_64 | ~ spl96_65 | ~ spl96_67 | ~ spl96_1009), inference(avatar_contradiction_clause, [], [f104267])).
fof(f104267, plain, ($false | (spl96_64 | ~ spl96_65 | ~ spl96_67 | ~ spl96_1009)), inference(subsumption_resolution, [], [f104266, f704])).
fof(f704, plain, (sP0(sK68) | ~ spl96_67), inference(avatar_component_clause, [], [f702])).
fof(f702, plain, (spl96_67 <=> sP0(sK68)), introduced(avatar_definition, [new_symbols(naming, [spl96_67])])).
fof(f104266, plain, (~ sP0(sK68) | (spl96_64 | ~ spl96_65 | ~ spl96_1009)), inference(resolution, [], [f101449, f696])).
fof(f696, plain, (r1(sK68, sK85) | ~ spl96_65), inference(avatar_component_clause, [], [f694])).
fof(f694, plain, (spl96_65 <=> r1(sK68, sK85)), introduced(avatar_definition, [new_symbols(naming, [spl96_65])])).
fof(f101449, plain, (! [X0] : (~ r1(X0, sK85) | ~ sP0(X0)) | (spl96_64 | ~ spl96_1009)), inference(subsumption_resolution, [], [f101448, f691])).
fof(f691, plain, (~ p2(sK85) | spl96_64), inference(avatar_component_clause, [], [f689])).
fof(f689, plain, (spl96_64 <=> p2(sK85)), introduced(avatar_definition, [new_symbols(naming, [spl96_64])])).
fof(f101448, plain, (! [X0] : (p2(sK85) | ~ r1(X0, sK85) | ~ sP0(X0)) | ~ spl96_1009), inference(resolution, [], [f53818, f312])).
fof(f312, plain, ! [X0, X1] : (~ p2(sK66(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ! [X0] : ((! [X1] : ((p2(sK65(X1)) & (~ p2(sK66(X1)) & r1(sK65(X1), sK66(X1))) & r1(X1, sK65(X1))) | p2(X1) | ~ r1(X0, X1)) & (~ p2(sK67(X0)) & r1(X0, sK67(X0)))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK65, sK66, sK67])], [f140, f143, f142, f141])).
fof(f141, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK65(X1)) & ? [X3] : (~ p2(X3) & r1(sK65(X1), X3)) & r1(X1, sK65(X1)))), introduced(choice_axiom, [])).
fof(f142, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK65(X1), X3)) => (~ p2(sK66(X1)) & r1(sK65(X1), sK66(X1)))), introduced(choice_axiom, [])).
fof(f143, plain, ! [X0] : (? [X4] : (~ p2(X4) & r1(X0, X4)) => (~ p2(sK67(X0)) & r1(X0, sK67(X0)))), introduced(choice_axiom, [])).
fof(f140, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (~ p2(X4) & r1(X0, X4))) | ~ sP0(X0)), inference(rectify, [], [f139])).
fof(f139, plain, ! [X0] : ((! [X135] : (? [X136] : (p2(X136) & ? [X137] : (~ p2(X137) & r1(X136, X137)) & r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) & ? [X138] : (~ p2(X138) & r1(X0, X138))) | ~ sP0(X0)), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ! [X0] : ((! [X135] : (? [X136] : (p2(X136) & ? [X137] : (~ p2(X137) & r1(X136, X137)) & r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) & ? [X138] : (~ p2(X138) & r1(X0, X138))) | ~ sP0(X0)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X0] : (sP0(X0) <=> (! [X135] : (? [X136] : (p2(X136) & ? [X137] : (~ p2(X137) & r1(X136, X137)) & r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) & ? [X138] : (~ p2(X138) & r1(X0, X138)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f53818, plain, (p2(sK66(sK85)) | ~ spl96_1009), inference(avatar_component_clause, [], [f53816])).
fof(f100894, plain, (~ spl96_67 | spl96_1117 | spl96_64 | ~ spl96_65), inference(avatar_split_clause, [], [f52659, f694, f689, f54342, f702])).
fof(f52659, plain, (p2(sK85) | r1(sK65(sK85), sK66(sK85)) | ~ sP0(sK68) | ~ spl96_65), inference(resolution, [], [f696, f311])).
fof(f311, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK65(X1), sK66(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f144])).
fof(f100893, plain, (~ spl96_67 | spl96_1118 | spl96_64 | ~ spl96_65), inference(avatar_split_clause, [], [f52658, f694, f689, f54348, f702])).
fof(f52658, plain, (p2(sK85) | r1(sK85, sK65(sK85)) | ~ sP0(sK68) | ~ spl96_65), inference(resolution, [], [f696, f310])).
fof(f310, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK65(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f144])).
fof(f100885, plain, (~ spl96_62 | spl96_385 | ~ spl96_4053), inference(avatar_contradiction_clause, [], [f100884])).
fof(f100884, plain, ($false | (~ spl96_62 | spl96_385 | ~ spl96_4053)), inference(subsumption_resolution, [], [f100883, f683])).
fof(f683, plain, (sP2(sK85) | ~ spl96_62), inference(avatar_component_clause, [], [f681])).
fof(f681, plain, (spl96_62 <=> sP2(sK85)), introduced(avatar_definition, [new_symbols(naming, [spl96_62])])).
fof(f100883, plain, (~ sP2(sK85) | (~ spl96_62 | spl96_385 | ~ spl96_4053)), inference(duplicate_literal_removal, [], [f100881])).
fof(f100881, plain, (~ sP2(sK85) | ~ sP2(sK85) | (~ spl96_62 | spl96_385 | ~ spl96_4053)), inference(resolution, [], [f100510, f296])).
fof(f296, plain, ! [X0] : (r1(X0, sK61(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ! [X0] : ((! [X1] : ((p2(sK59(X1)) & (~ p2(sK60(X1)) & r1(sK59(X1), sK60(X1))) & r1(X1, sK59(X1))) | p2(X1) | ~ r1(X0, X1)) & ((! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK62(X0), X6)) & ~ p2(sK62(X0)) & r1(sK61(X0), sK62(X0))) & r1(X0, sK61(X0)))) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK59, sK60, sK61, sK62])], [f128, f132, f131, f130, f129])).
fof(f129, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK59(X1)) & ? [X3] : (~ p2(X3) & r1(sK59(X1), X3)) & r1(X1, sK59(X1)))), introduced(choice_axiom, [])).
fof(f130, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK59(X1), X3)) => (~ p2(sK60(X1)) & r1(sK59(X1), sK60(X1)))), introduced(choice_axiom, [])).
fof(f131, plain, ! [X0] : (? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4)) => (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK61(X0), X5)) & r1(X0, sK61(X0)))), introduced(choice_axiom, [])).
fof(f132, plain, ! [X0] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK61(X0), X5)) => (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK62(X0), X6)) & ~ p2(sK62(X0)) & r1(sK61(X0), sK62(X0)))), introduced(choice_axiom, [])).
fof(f128, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4))) | ~ sP2(X0)), inference(rectify, [], [f127])).
fof(f127, plain, ! [X92] : ((! [X115] : (? [X116] : (p2(X116) & ? [X117] : (~ p2(X117) & r1(X116, X117)) & r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) & ? [X118] : (? [X119] : (! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) & ~ p2(X119) & r1(X118, X119)) & r1(X92, X118))) | ~ sP2(X92)), inference(nnf_transformation, [], [f11])).
fof(f11, plain, ! [X92] : ((! [X115] : (? [X116] : (p2(X116) & ? [X117] : (~ p2(X117) & r1(X116, X117)) & r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) & ? [X118] : (? [X119] : (! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) & ~ p2(X119) & r1(X118, X119)) & r1(X92, X118))) | ~ sP2(X92)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X92] : (sP2(X92) <=> (! [X115] : (? [X116] : (p2(X116) & ? [X117] : (~ p2(X117) & r1(X116, X117)) & r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) & ? [X118] : (? [X119] : (! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) & ~ p2(X119) & r1(X118, X119)) & r1(X92, X118)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f100510, plain, (! [X0] : (~ r1(X0, sK61(sK85)) | ~ sP2(X0)) | (~ spl96_62 | spl96_385 | ~ spl96_4053)), inference(subsumption_resolution, [], [f100509, f29745])).
fof(f29745, plain, (~ p2(sK61(sK85)) | spl96_385), inference(avatar_component_clause, [], [f29743])).
fof(f29743, plain, (spl96_385 <=> p2(sK61(sK85))), introduced(avatar_definition, [new_symbols(naming, [spl96_385])])).
fof(f100509, plain, (! [X0] : (p2(sK61(sK85)) | ~ r1(X0, sK61(sK85)) | ~ sP2(X0)) | (~ spl96_62 | spl96_385 | ~ spl96_4053)), inference(resolution, [], [f98544, f302])).
fof(f302, plain, ! [X0, X1] : (~ p2(sK60(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP2(X0)), inference(cnf_transformation, [], [f133])).
fof(f98544, plain, (p2(sK60(sK61(sK85))) | (~ spl96_62 | spl96_385 | ~ spl96_4053)), inference(subsumption_resolution, [], [f98543, f683])).
fof(f98543, plain, (p2(sK60(sK61(sK85))) | ~ sP2(sK85) | (spl96_385 | ~ spl96_4053)), inference(subsumption_resolution, [], [f98514, f29745])).
fof(f98514, plain, (p2(sK60(sK61(sK85))) | p2(sK61(sK85)) | ~ sP2(sK85) | ~ spl96_4053), inference(resolution, [], [f97838, f1920])).
fof(f1920, plain, ! [X13] : (r1(sK59(sK61(X13)), sK60(sK61(X13))) | p2(sK61(X13)) | ~ sP2(X13)), inference(duplicate_literal_removal, [], [f1897])).
fof(f1897, plain, ! [X13] : (p2(sK61(X13)) | r1(sK59(sK61(X13)), sK60(sK61(X13))) | ~ sP2(X13) | ~ sP2(X13)), inference(resolution, [], [f301, f296])).
fof(f301, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK59(X1), sK60(X1)) | ~ sP2(X0)), inference(cnf_transformation, [], [f133])).
fof(f97838, plain, (! [X3] : (~ r1(sK59(sK61(sK85)), X3) | p2(X3)) | ~ spl96_4053), inference(avatar_component_clause, [], [f97837])).
fof(f97837, plain, (spl96_4053 <=> ! [X3] : (p2(X3) | ~ r1(sK59(sK61(sK85)), X3))), introduced(avatar_definition, [new_symbols(naming, [spl96_4053])])).
fof(f98165, plain, (~ spl96_62 | spl96_385 | spl96_4052), inference(avatar_contradiction_clause, [], [f98164])).
fof(f98164, plain, ($false | (~ spl96_62 | spl96_385 | spl96_4052)), inference(subsumption_resolution, [], [f98163, f683])).
fof(f98163, plain, (~ sP2(sK85) | (spl96_385 | spl96_4052)), inference(subsumption_resolution, [], [f98161, f29745])).
fof(f98161, plain, (p2(sK61(sK85)) | ~ sP2(sK85) | spl96_4052), inference(resolution, [], [f97835, f1438])).
fof(f1438, plain, ! [X13] : (p2(sK59(sK61(X13))) | p2(sK61(X13)) | ~ sP2(X13)), inference(duplicate_literal_removal, [], [f1415])).
fof(f1415, plain, ! [X13] : (p2(sK61(X13)) | p2(sK59(sK61(X13))) | ~ sP2(X13) | ~ sP2(X13)), inference(resolution, [], [f303, f296])).
fof(f303, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK59(X1)) | ~ sP2(X0)), inference(cnf_transformation, [], [f133])).
fof(f97835, plain, (~ p2(sK59(sK61(sK85))) | spl96_4052), inference(avatar_component_clause, [], [f97833])).
fof(f97833, plain, (spl96_4052 <=> p2(sK59(sK61(sK85)))), introduced(avatar_definition, [new_symbols(naming, [spl96_4052])])).
fof(f97839, plain, (~ spl96_4052 | spl96_4053 | ~ spl96_62 | spl96_385 | ~ spl96_3952), inference(avatar_split_clause, [], [f97831, f97136, f29743, f681, f97837, f97833])).
fof(f97136, plain, (spl96_3952 <=> ! [X38, X39] : (~ p2(X38) | p2(X39) | ~ r1(X38, X39) | ~ r1(sK61(sK85), X38))), introduced(avatar_definition, [new_symbols(naming, [spl96_3952])])).
fof(f97831, plain, (! [X3] : (p2(X3) | ~ r1(sK59(sK61(sK85)), X3) | ~ p2(sK59(sK61(sK85)))) | (~ spl96_62 | spl96_385 | ~ spl96_3952)), inference(subsumption_resolution, [], [f97830, f683])).
fof(f97830, plain, (! [X3] : (p2(X3) | ~ r1(sK59(sK61(sK85)), X3) | ~ p2(sK59(sK61(sK85))) | ~ sP2(sK85)) | (spl96_385 | ~ spl96_3952)), inference(subsumption_resolution, [], [f97789, f29745])).
fof(f97789, plain, (! [X3] : (p2(X3) | ~ r1(sK59(sK61(sK85)), X3) | ~ p2(sK59(sK61(sK85))) | p2(sK61(sK85)) | ~ sP2(sK85)) | ~ spl96_3952), inference(resolution, [], [f97137, f1697])).
fof(f1697, plain, ! [X13] : (r1(sK61(X13), sK59(sK61(X13))) | p2(sK61(X13)) | ~ sP2(X13)), inference(duplicate_literal_removal, [], [f1674])).
fof(f1674, plain, ! [X13] : (p2(sK61(X13)) | r1(sK61(X13), sK59(sK61(X13))) | ~ sP2(X13) | ~ sP2(X13)), inference(resolution, [], [f300, f296])).
fof(f300, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK59(X1)) | ~ sP2(X0)), inference(cnf_transformation, [], [f133])).
fof(f97137, plain, (! [X39, X38] : (~ r1(sK61(sK85), X38) | p2(X39) | ~ r1(X38, X39) | ~ p2(X38)) | ~ spl96_3952), inference(avatar_component_clause, [], [f97136])).
fof(f97745, plain, (spl96_64 | spl96_815 | ~ spl96_63 | ~ spl96_65 | ~ spl96_73 | ~ spl96_816), inference(avatar_split_clause, [], [f97744, f52237, f726, f694, f685, f52234, f689])).
fof(f52234, plain, (spl96_815 <=> ! [X26] : (~ r1(sK93(sK85), X26) | p2(X26))), introduced(avatar_definition, [new_symbols(naming, [spl96_815])])).
fof(f726, plain, (spl96_73 <=> ! [X57] : (r1(X57, sK93(X57)) | ~ r1(sK68, X57) | p2(X57))), introduced(avatar_definition, [new_symbols(naming, [spl96_73])])).
fof(f52237, plain, (spl96_816 <=> p2(sK93(sK85))), introduced(avatar_definition, [new_symbols(naming, [spl96_816])])).
fof(f97744, plain, (! [X27] : (~ r1(sK93(sK85), X27) | p2(X27) | p2(sK85)) | (~ spl96_63 | ~ spl96_65 | ~ spl96_73 | ~ spl96_816)), inference(subsumption_resolution, [], [f97743, f696])).
fof(f97743, plain, (! [X27] : (~ r1(sK93(sK85), X27) | p2(X27) | ~ r1(sK68, sK85) | p2(sK85)) | (~ spl96_63 | ~ spl96_73 | ~ spl96_816)), inference(subsumption_resolution, [], [f97446, f52238])).
fof(f52238, plain, (p2(sK93(sK85)) | ~ spl96_816), inference(avatar_component_clause, [], [f52237])).
fof(f97446, plain, (! [X27] : (~ p2(sK93(sK85)) | ~ r1(sK93(sK85), X27) | p2(X27) | ~ r1(sK68, sK85) | p2(sK85)) | (~ spl96_63 | ~ spl96_73)), inference(resolution, [], [f686, f727])).
fof(f727, plain, (! [X57] : (r1(X57, sK93(X57)) | ~ r1(sK68, X57) | p2(X57)) | ~ spl96_73), inference(avatar_component_clause, [], [f726])).
fof(f97742, plain, (~ spl96_62 | spl96_3952 | ~ spl96_60), inference(avatar_split_clause, [], [f97741, f673, f97136, f681])).
fof(f673, plain, (spl96_60 <=> ! [X42, X44, X43] : (~ p2(X43) | ~ r1(sK85, X42) | sP5(X42) | sP4(X42) | ~ r1(X42, X43) | ~ r1(X43, X44) | p2(X44))), introduced(avatar_definition, [new_symbols(naming, [spl96_60])])).
fof(f97741, plain, (! [X39, X38] : (~ p2(X38) | ~ r1(sK61(sK85), X38) | ~ r1(X38, X39) | p2(X39) | ~ sP2(sK85)) | ~ spl96_60), inference(subsumption_resolution, [], [f97740, f33134])).
fof(f33134, plain, ! [X1] : (~ sP4(sK61(X1)) | ~ sP2(X1)), inference(duplicate_literal_removal, [], [f33133])).
fof(f33133, plain, ! [X1] : (~ sP4(sK61(X1)) | ~ sP2(X1) | ~ sP4(sK61(X1)) | ~ sP2(X1)), inference(resolution, [], [f32002, f297])).
fof(f297, plain, ! [X0] : (r1(sK61(X0), sK62(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f133])).
fof(f32002, plain, ! [X0, X1] : (~ r1(X1, sK62(X0)) | ~ sP4(sK61(X0)) | ~ sP2(X0) | ~ sP4(X1)), inference(subsumption_resolution, [], [f32001, f298])).
fof(f298, plain, ! [X0] : (~ p2(sK62(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f133])).
fof(f32001, plain, ! [X0, X1] : (~ sP2(X0) | ~ sP4(sK61(X0)) | p2(sK62(X0)) | ~ r1(X1, sK62(X0)) | ~ sP4(X1)), inference(resolution, [], [f31866, f290])).
fof(f290, plain, ! [X0, X1] : (~ p2(sK54(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP4(X0)), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ! [X0] : ((! [X1] : ((p2(sK53(X1)) & (~ p2(sK54(X1)) & r1(sK53(X1), sK54(X1))) & r1(X1, sK53(X1))) | p2(X1) | ~ r1(X0, X1)) & ((! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK56(X0), X6)) & ~ p2(sK56(X0)) & r1(sK55(X0), sK56(X0))) & r1(X0, sK55(X0)))) | ~ sP4(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK53, sK54, sK55, sK56])], [f116, f120, f119, f118, f117])).
fof(f117, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK53(X1)) & ? [X3] : (~ p2(X3) & r1(sK53(X1), X3)) & r1(X1, sK53(X1)))), introduced(choice_axiom, [])).
fof(f118, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK53(X1), X3)) => (~ p2(sK54(X1)) & r1(sK53(X1), sK54(X1)))), introduced(choice_axiom, [])).
fof(f119, plain, ! [X0] : (? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4)) => (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK55(X0), X5)) & r1(X0, sK55(X0)))), introduced(choice_axiom, [])).
fof(f120, plain, ! [X0] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK55(X0), X5)) => (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK56(X0), X6)) & ~ p2(sK56(X0)) & r1(sK55(X0), sK56(X0)))), introduced(choice_axiom, [])).
fof(f116, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4))) | ~ sP4(X0)), inference(rectify, [], [f115])).
fof(f115, plain, ! [X93] : ((! [X96] : (? [X97] : (p2(X97) & ? [X98] : (~ p2(X98) & r1(X97, X98)) & r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) & ? [X99] : (? [X100] : (! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) & ~ p2(X100) & r1(X99, X100)) & r1(X93, X99))) | ~ sP4(X93)), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ! [X93] : ((! [X96] : (? [X97] : (p2(X97) & ? [X98] : (~ p2(X98) & r1(X97, X98)) & r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) & ? [X99] : (? [X100] : (! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) & ~ p2(X100) & r1(X99, X100)) & r1(X93, X99))) | ~ sP4(X93)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X93] : (sP4(X93) <=> (! [X96] : (? [X97] : (p2(X97) & ? [X98] : (~ p2(X98) & r1(X97, X98)) & r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) & ? [X99] : (? [X100] : (! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) & ~ p2(X100) & r1(X99, X100)) & r1(X93, X99)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f31866, plain, ! [X0] : (p2(sK54(sK62(X0))) | ~ sP2(X0) | ~ sP4(sK61(X0))), inference(duplicate_literal_removal, [], [f31837])).
fof(f31837, plain, ! [X0] : (~ sP2(X0) | p2(sK54(sK62(X0))) | ~ sP4(sK61(X0)) | ~ sP4(sK61(X0)) | ~ sP2(X0)), inference(resolution, [], [f7122, f1882])).
fof(f1882, plain, ! [X30] : (r1(sK53(sK62(X30)), sK54(sK62(X30))) | ~ sP4(sK61(X30)) | ~ sP2(X30)), inference(subsumption_resolution, [], [f1861, f298])).
fof(f1861, plain, ! [X30] : (p2(sK62(X30)) | r1(sK53(sK62(X30)), sK54(sK62(X30))) | ~ sP4(sK61(X30)) | ~ sP2(X30)), inference(resolution, [], [f289, f297])).
fof(f289, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK53(X1), sK54(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f121])).
fof(f7122, plain, ! [X0, X1] : (~ r1(sK53(sK62(X0)), X1) | ~ sP2(X0) | p2(X1) | ~ sP4(sK61(X0))), inference(subsumption_resolution, [], [f7121, f1400])).
fof(f1400, plain, ! [X30] : (p2(sK53(sK62(X30))) | ~ sP4(sK61(X30)) | ~ sP2(X30)), inference(subsumption_resolution, [], [f1379, f298])).
fof(f1379, plain, ! [X30] : (p2(sK62(X30)) | p2(sK53(sK62(X30))) | ~ sP4(sK61(X30)) | ~ sP2(X30)), inference(resolution, [], [f291, f297])).
fof(f291, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK53(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f121])).
fof(f7121, plain, ! [X0, X1] : (~ sP4(sK61(X0)) | ~ sP2(X0) | p2(X1) | ~ r1(sK53(sK62(X0)), X1) | ~ p2(sK53(sK62(X0)))), inference(duplicate_literal_removal, [], [f7071])).
fof(f7071, plain, ! [X0, X1] : (~ sP4(sK61(X0)) | ~ sP2(X0) | p2(X1) | ~ r1(sK53(sK62(X0)), X1) | ~ p2(sK53(sK62(X0))) | ~ sP2(X0)), inference(resolution, [], [f1659, f299])).
fof(f299, plain, ! [X6, X0, X7] : (~ r1(sK62(X0), X6) | p2(X7) | ~ r1(X6, X7) | ~ p2(X6) | ~ sP2(X0)), inference(cnf_transformation, [], [f133])).
fof(f1659, plain, ! [X30] : (r1(sK62(X30), sK53(sK62(X30))) | ~ sP4(sK61(X30)) | ~ sP2(X30)), inference(subsumption_resolution, [], [f1638, f298])).
fof(f1638, plain, ! [X30] : (p2(sK62(X30)) | r1(sK62(X30), sK53(sK62(X30))) | ~ sP4(sK61(X30)) | ~ sP2(X30)), inference(resolution, [], [f288, f297])).
fof(f288, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK53(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f121])).
fof(f97740, plain, (! [X39, X38] : (~ p2(X38) | sP4(sK61(sK85)) | ~ r1(sK61(sK85), X38) | ~ r1(X38, X39) | p2(X39) | ~ sP2(sK85)) | ~ spl96_60), inference(subsumption_resolution, [], [f97115, f25444])).
fof(f25444, plain, ! [X1] : (~ sP5(sK61(X1)) | ~ sP2(X1)), inference(duplicate_literal_removal, [], [f25443])).
fof(f25443, plain, ! [X1] : (~ sP5(sK61(X1)) | ~ sP2(X1) | ~ sP5(sK61(X1)) | ~ sP2(X1)), inference(resolution, [], [f25005, f297])).
fof(f25005, plain, ! [X0, X1] : (~ r1(X1, sK62(X0)) | ~ sP5(sK61(X0)) | ~ sP2(X0) | ~ sP5(X1)), inference(subsumption_resolution, [], [f25004, f298])).
fof(f25004, plain, ! [X0, X1] : (~ sP2(X0) | ~ sP5(sK61(X0)) | p2(sK62(X0)) | ~ r1(X1, sK62(X0)) | ~ sP5(X1)), inference(resolution, [], [f24992, f282])).
fof(f282, plain, ! [X0, X1] : (~ p2(sK51(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP5(X0)), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ! [X0] : (! [X1] : ((((p2(sK50(X1)) & (~ p2(sK51(X1)) & r1(sK50(X1), sK51(X1))) & r1(X1, sK50(X1))) | p2(X1)) & ((! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(sK52(X1), X5)) & ~ p2(sK52(X1)) & r1(X1, sK52(X1))) | sP3(X1))) | ~ r1(X0, X1)) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK50, sK51, sK52])], [f110, f113, f112, f111])).
fof(f111, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK50(X1)) & ? [X3] : (~ p2(X3) & r1(sK50(X1), X3)) & r1(X1, sK50(X1)))), introduced(choice_axiom, [])).
fof(f112, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK50(X1), X3)) => (~ p2(sK51(X1)) & r1(sK50(X1), sK51(X1)))), introduced(choice_axiom, [])).
fof(f113, plain, ! [X1] : (? [X4] : (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(X4, X5)) & ~ p2(X4) & r1(X1, X4)) => (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(sK52(X1), X5)) & ~ p2(sK52(X1)) & r1(X1, sK52(X1)))), introduced(choice_axiom, [])).
fof(f110, plain, ! [X0] : (! [X1] : (((? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1)) & (? [X4] : (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(X4, X5)) & ~ p2(X4) & r1(X1, X4)) | sP3(X1))) | ~ r1(X0, X1)) | ~ sP5(X0)), inference(rectify, [], [f109])).
fof(f109, plain, ! [X93] : (! [X103] : (((? [X104] : (p2(X104) & ? [X105] : (~ p2(X105) & r1(X104, X105)) & r1(X103, X104)) | p2(X103)) & (? [X106] : (! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) & ~ p2(X106) & r1(X103, X106)) | sP3(X103))) | ~ r1(X93, X103)) | ~ sP5(X93)), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ! [X93] : (! [X103] : (((? [X104] : (p2(X104) & ? [X105] : (~ p2(X105) & r1(X104, X105)) & r1(X103, X104)) | p2(X103)) & (? [X106] : (! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) & ~ p2(X106) & r1(X103, X106)) | sP3(X103))) | ~ r1(X93, X103)) | ~ sP5(X93)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X93] : (sP5(X93) <=> ! [X103] : (((? [X104] : (p2(X104) & ? [X105] : (~ p2(X105) & r1(X104, X105)) & r1(X103, X104)) | p2(X103)) & (? [X106] : (! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) & ~ p2(X106) & r1(X103, X106)) | sP3(X103))) | ~ r1(X93, X103))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f24992, plain, ! [X0] : (p2(sK51(sK62(X0))) | ~ sP2(X0) | ~ sP5(sK61(X0))), inference(duplicate_literal_removal, [], [f24963])).
fof(f24963, plain, ! [X0] : (~ sP2(X0) | p2(sK51(sK62(X0))) | ~ sP5(sK61(X0)) | ~ sP5(sK61(X0)) | ~ sP2(X0)), inference(resolution, [], [f6941, f1829])).
fof(f1829, plain, ! [X30] : (r1(sK50(sK62(X30)), sK51(sK62(X30))) | ~ sP5(sK61(X30)) | ~ sP2(X30)), inference(subsumption_resolution, [], [f1809, f298])).
fof(f1809, plain, ! [X30] : (p2(sK62(X30)) | r1(sK50(sK62(X30)), sK51(sK62(X30))) | ~ sP5(sK61(X30)) | ~ sP2(X30)), inference(resolution, [], [f281, f297])).
fof(f281, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK50(X1), sK51(X1)) | ~ sP5(X0)), inference(cnf_transformation, [], [f114])).
fof(f6941, plain, ! [X0, X1] : (~ r1(sK50(sK62(X0)), X1) | ~ sP2(X0) | p2(X1) | ~ sP5(sK61(X0))), inference(subsumption_resolution, [], [f6940, f1333])).
fof(f1333, plain, ! [X30] : (p2(sK50(sK62(X30))) | ~ sP5(sK61(X30)) | ~ sP2(X30)), inference(subsumption_resolution, [], [f1313, f298])).
fof(f1313, plain, ! [X30] : (p2(sK62(X30)) | p2(sK50(sK62(X30))) | ~ sP5(sK61(X30)) | ~ sP2(X30)), inference(resolution, [], [f283, f297])).
fof(f283, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK50(X1)) | ~ sP5(X0)), inference(cnf_transformation, [], [f114])).
fof(f6940, plain, ! [X0, X1] : (~ sP5(sK61(X0)) | ~ sP2(X0) | p2(X1) | ~ r1(sK50(sK62(X0)), X1) | ~ p2(sK50(sK62(X0)))), inference(duplicate_literal_removal, [], [f6890])).
fof(f6890, plain, ! [X0, X1] : (~ sP5(sK61(X0)) | ~ sP2(X0) | p2(X1) | ~ r1(sK50(sK62(X0)), X1) | ~ p2(sK50(sK62(X0))) | ~ sP2(X0)), inference(resolution, [], [f1606, f299])).
fof(f1606, plain, ! [X30] : (r1(sK62(X30), sK50(sK62(X30))) | ~ sP5(sK61(X30)) | ~ sP2(X30)), inference(subsumption_resolution, [], [f1586, f298])).
fof(f1586, plain, ! [X30] : (p2(sK62(X30)) | r1(sK62(X30), sK50(sK62(X30))) | ~ sP5(sK61(X30)) | ~ sP2(X30)), inference(resolution, [], [f280, f297])).
fof(f280, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK50(X1)) | ~ sP5(X0)), inference(cnf_transformation, [], [f114])).
fof(f97115, plain, (! [X39, X38] : (~ p2(X38) | sP5(sK61(sK85)) | sP4(sK61(sK85)) | ~ r1(sK61(sK85), X38) | ~ r1(X38, X39) | p2(X39) | ~ sP2(sK85)) | ~ spl96_60), inference(resolution, [], [f674, f296])).
fof(f674, plain, (! [X43, X44, X42] : (~ r1(sK85, X42) | ~ p2(X43) | sP5(X42) | sP4(X42) | ~ r1(X42, X43) | ~ r1(X43, X44) | p2(X44)) | ~ spl96_60), inference(avatar_component_clause, [], [f673])).
fof(f97708, plain, (spl96_64 | ~ spl96_65 | ~ spl96_71 | ~ spl96_72 | ~ spl96_815), inference(avatar_contradiction_clause, [], [f97707])).
fof(f97707, plain, ($false | (spl96_64 | ~ spl96_65 | ~ spl96_71 | ~ spl96_72 | ~ spl96_815)), inference(subsumption_resolution, [], [f97706, f691])).
fof(f97706, plain, (p2(sK85) | (spl96_64 | ~ spl96_65 | ~ spl96_71 | ~ spl96_72 | ~ spl96_815)), inference(subsumption_resolution, [], [f97705, f696])).
fof(f97705, plain, (~ r1(sK68, sK85) | p2(sK85) | (spl96_64 | ~ spl96_65 | ~ spl96_71 | ~ spl96_72 | ~ spl96_815)), inference(resolution, [], [f97491, f719])).
fof(f719, plain, (! [X57] : (~ p2(sK94(X57)) | ~ r1(sK68, X57) | p2(X57)) | ~ spl96_71), inference(avatar_component_clause, [], [f718])).
fof(f718, plain, (spl96_71 <=> ! [X57] : (~ p2(sK94(X57)) | ~ r1(sK68, X57) | p2(X57))), introduced(avatar_definition, [new_symbols(naming, [spl96_71])])).
fof(f97491, plain, (p2(sK94(sK85)) | (spl96_64 | ~ spl96_65 | ~ spl96_72 | ~ spl96_815)), inference(subsumption_resolution, [], [f97490, f691])).
fof(f97490, plain, (p2(sK94(sK85)) | p2(sK85) | (~ spl96_65 | ~ spl96_72 | ~ spl96_815)), inference(subsumption_resolution, [], [f97450, f696])).
fof(f97450, plain, (p2(sK94(sK85)) | ~ r1(sK68, sK85) | p2(sK85) | (~ spl96_72 | ~ spl96_815)), inference(resolution, [], [f52235, f723])).
fof(f723, plain, (! [X57] : (r1(sK93(X57), sK94(X57)) | ~ r1(sK68, X57) | p2(X57)) | ~ spl96_72), inference(avatar_component_clause, [], [f722])).
fof(f722, plain, (spl96_72 <=> ! [X57] : (r1(sK93(X57), sK94(X57)) | ~ r1(sK68, X57) | p2(X57))), introduced(avatar_definition, [new_symbols(naming, [spl96_72])])).
fof(f52235, plain, (! [X26] : (~ r1(sK93(sK85), X26) | p2(X26)) | ~ spl96_815), inference(avatar_component_clause, [], [f52234])).
fof(f97418, plain, (spl96_816 | spl96_64 | ~ spl96_65 | ~ spl96_70), inference(avatar_split_clause, [], [f97417, f714, f694, f689, f52237])).
fof(f714, plain, (spl96_70 <=> ! [X57] : (p2(sK93(X57)) | ~ r1(sK68, X57) | p2(X57))), introduced(avatar_definition, [new_symbols(naming, [spl96_70])])).
fof(f97417, plain, (p2(sK93(sK85)) | (spl96_64 | ~ spl96_65 | ~ spl96_70)), inference(subsumption_resolution, [], [f63796, f691])).
fof(f63796, plain, (p2(sK93(sK85)) | p2(sK85) | (~ spl96_65 | ~ spl96_70)), inference(resolution, [], [f715, f696])).
fof(f715, plain, (! [X57] : (~ r1(sK68, X57) | p2(sK93(X57)) | p2(X57)) | ~ spl96_70), inference(avatar_component_clause, [], [f714])).
fof(f97391, plain, (~ spl96_62 | ~ spl96_384), inference(avatar_split_clause, [], [f97388, f29739, f681])).
fof(f29739, plain, (spl96_384 <=> sP4(sK61(sK85))), introduced(avatar_definition, [new_symbols(naming, [spl96_384])])).
fof(f97388, plain, (~ sP2(sK85) | ~ spl96_384), inference(resolution, [], [f29741, f33134])).
fof(f29741, plain, (sP4(sK61(sK85)) | ~ spl96_384), inference(avatar_component_clause, [], [f29739])).
fof(f97089, plain, (~ spl96_59 | spl96_140 | ~ spl96_70 | ~ spl96_71 | ~ spl96_72 | ~ spl96_73), inference(avatar_split_clause, [], [f95561, f726, f722, f718, f714, f1166, f669])).
fof(f669, plain, (spl96_59 <=> sP6(sK68)), introduced(avatar_definition, [new_symbols(naming, [spl96_59])])).
fof(f1166, plain, (spl96_140 <=> sP1(sK68)), introduced(avatar_definition, [new_symbols(naming, [spl96_140])])).
fof(f95561, plain, (sP1(sK68) | ~ sP6(sK68) | (~ spl96_70 | ~ spl96_71 | ~ spl96_72 | ~ spl96_73)), inference(duplicate_literal_removal, [], [f95560])).
fof(f95560, plain, (sP1(sK68) | ~ sP6(sK68) | sP1(sK68) | ~ sP6(sK68) | (~ spl96_70 | ~ spl96_71 | ~ spl96_72 | ~ spl96_73)), inference(resolution, [], [f93193, f270])).
fof(f270, plain, ! [X0] : (r1(X0, sK49(X0)) | sP1(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ! [X0] : ((((p2(sK47(X0)) & (~ p2(sK48(X0)) & r1(sK47(X0), sK48(X0))) & r1(X0, sK47(X0))) | p2(X0)) & ((! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(sK49(X0), X4)) & ~ p2(sK49(X0)) & r1(X0, sK49(X0))) | sP1(X0))) | ~ sP6(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK47, sK48, sK49])], [f104, f107, f106, f105])).
fof(f105, plain, ! [X0] : (? [X1] : (p2(X1) & ? [X2] : (~ p2(X2) & r1(X1, X2)) & r1(X0, X1)) => (p2(sK47(X0)) & ? [X2] : (~ p2(X2) & r1(sK47(X0), X2)) & r1(X0, sK47(X0)))), introduced(choice_axiom, [])).
fof(f106, plain, ! [X0] : (? [X2] : (~ p2(X2) & r1(sK47(X0), X2)) => (~ p2(sK48(X0)) & r1(sK47(X0), sK48(X0)))), introduced(choice_axiom, [])).
fof(f107, plain, ! [X0] : (? [X3] : (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) & ~ p2(X3) & r1(X0, X3)) => (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(sK49(X0), X4)) & ~ p2(sK49(X0)) & r1(X0, sK49(X0)))), introduced(choice_axiom, [])).
fof(f104, plain, ! [X0] : (((? [X1] : (p2(X1) & ? [X2] : (~ p2(X2) & r1(X1, X2)) & r1(X0, X1)) | p2(X0)) & (? [X3] : (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) & ~ p2(X3) & r1(X0, X3)) | sP1(X0))) | ~ sP6(X0)), inference(rectify, [], [f103])).
fof(f103, plain, ! [X0] : (((? [X122] : (p2(X122) & ? [X123] : (~ p2(X123) & r1(X122, X123)) & r1(X0, X122)) | p2(X0)) & (? [X124] : (! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) & ~ p2(X124) & r1(X0, X124)) | sP1(X0))) | ~ sP6(X0)), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ! [X0] : (((? [X122] : (p2(X122) & ? [X123] : (~ p2(X123) & r1(X122, X123)) & r1(X0, X122)) | p2(X0)) & (? [X124] : (! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) & ~ p2(X124) & r1(X0, X124)) | sP1(X0))) | ~ sP6(X0)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X0] : (sP6(X0) <=> ((? [X122] : (p2(X122) & ? [X123] : (~ p2(X123) & r1(X122, X123)) & r1(X0, X122)) | p2(X0)) & (? [X124] : (! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) & ~ p2(X124) & r1(X0, X124)) | sP1(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f93193, plain, (! [X1] : (~ r1(sK68, sK49(X1)) | sP1(X1) | ~ sP6(X1)) | (~ spl96_70 | ~ spl96_71 | ~ spl96_72 | ~ spl96_73)), inference(subsumption_resolution, [], [f93192, f271])).
fof(f271, plain, ! [X0] : (~ p2(sK49(X0)) | sP1(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f108])).
fof(f93192, plain, (! [X1] : (~ r1(sK68, sK49(X1)) | sP1(X1) | ~ sP6(X1) | p2(sK49(X1))) | (~ spl96_70 | ~ spl96_71 | ~ spl96_72 | ~ spl96_73)), inference(subsumption_resolution, [], [f93191, f715])).
fof(f93191, plain, (! [X1] : (~ r1(sK68, sK49(X1)) | ~ p2(sK93(sK49(X1))) | sP1(X1) | ~ sP6(X1) | p2(sK49(X1))) | (~ spl96_71 | ~ spl96_72 | ~ spl96_73)), inference(subsumption_resolution, [], [f93187, f719])).
fof(f93187, plain, (! [X1] : (p2(sK94(sK49(X1))) | ~ r1(sK68, sK49(X1)) | ~ p2(sK93(sK49(X1))) | sP1(X1) | ~ sP6(X1) | p2(sK49(X1))) | (~ spl96_72 | ~ spl96_73)), inference(duplicate_literal_removal, [], [f93159])).
fof(f93159, plain, (! [X1] : (p2(sK94(sK49(X1))) | ~ r1(sK68, sK49(X1)) | ~ p2(sK93(sK49(X1))) | sP1(X1) | ~ sP6(X1) | ~ r1(sK68, sK49(X1)) | p2(sK49(X1))) | (~ spl96_72 | ~ spl96_73)), inference(resolution, [], [f91512, f723])).
fof(f91512, plain, (! [X306, X305] : (~ r1(sK93(sK49(X305)), X306) | p2(X306) | ~ r1(sK68, sK49(X305)) | ~ p2(sK93(sK49(X305))) | sP1(X305) | ~ sP6(X305)) | ~ spl96_73), inference(subsumption_resolution, [], [f91399, f271])).
fof(f91399, plain, (! [X306, X305] : (~ r1(sK68, sK49(X305)) | p2(sK49(X305)) | p2(X306) | ~ r1(sK93(sK49(X305)), X306) | ~ p2(sK93(sK49(X305))) | sP1(X305) | ~ sP6(X305)) | ~ spl96_73), inference(resolution, [], [f727, f272])).
fof(f272, plain, ! [X4, X0, X5] : (~ r1(sK49(X0), X4) | p2(X5) | ~ r1(X4, X5) | ~ p2(X4) | sP1(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f108])).
fof(f97072, plain, (spl96_3799 | spl96_74 | ~ spl96_75 | ~ spl96_1365 | ~ spl96_1369 | ~ spl96_1371), inference(avatar_split_clause, [], [f95579, f63756, f63744, f63722, f735, f730, f94883])).
fof(f94883, plain, (spl96_3799 <=> p2(sK64(sK69(sK95)))), introduced(avatar_definition, [new_symbols(naming, [spl96_3799])])).
fof(f730, plain, (spl96_74 <=> p2(sK95)), introduced(avatar_definition, [new_symbols(naming, [spl96_74])])).
fof(f735, plain, (spl96_75 <=> r1(sK68, sK95)), introduced(avatar_definition, [new_symbols(naming, [spl96_75])])).
fof(f63722, plain, (spl96_1365 <=> p2(sK63(sK69(sK95)))), introduced(avatar_definition, [new_symbols(naming, [spl96_1365])])).
fof(f63744, plain, (spl96_1369 <=> r1(sK69(sK95), sK63(sK69(sK95)))), introduced(avatar_definition, [new_symbols(naming, [spl96_1369])])).
fof(f63756, plain, (spl96_1371 <=> r1(sK63(sK69(sK95)), sK64(sK69(sK95)))), introduced(avatar_definition, [new_symbols(naming, [spl96_1371])])).
fof(f95579, plain, (p2(sK64(sK69(sK95))) | (spl96_74 | ~ spl96_75 | ~ spl96_1365 | ~ spl96_1369 | ~ spl96_1371)), inference(resolution, [], [f94190, f63758])).
fof(f63758, plain, (r1(sK63(sK69(sK95)), sK64(sK69(sK95))) | ~ spl96_1371), inference(avatar_component_clause, [], [f63756])).
fof(f94190, plain, (! [X0] : (~ r1(sK63(sK69(sK95)), X0) | p2(X0)) | (spl96_74 | ~ spl96_75 | ~ spl96_1365 | ~ spl96_1369)), inference(subsumption_resolution, [], [f94189, f737])).
fof(f737, plain, (r1(sK68, sK95) | ~ spl96_75), inference(avatar_component_clause, [], [f735])).
fof(f94189, plain, (! [X0] : (p2(X0) | ~ r1(sK63(sK69(sK95)), X0) | ~ r1(sK68, sK95)) | (spl96_74 | ~ spl96_1365 | ~ spl96_1369)), inference(subsumption_resolution, [], [f94188, f732])).
fof(f732, plain, (~ p2(sK95) | spl96_74), inference(avatar_component_clause, [], [f730])).
fof(f94188, plain, (! [X0] : (p2(X0) | ~ r1(sK63(sK69(sK95)), X0) | p2(sK95) | ~ r1(sK68, sK95)) | (~ spl96_1365 | ~ spl96_1369)), inference(subsumption_resolution, [], [f94028, f63724])).
fof(f63724, plain, (p2(sK63(sK69(sK95))) | ~ spl96_1365), inference(avatar_component_clause, [], [f63722])).
fof(f94028, plain, (! [X0] : (p2(X0) | ~ r1(sK63(sK69(sK95)), X0) | ~ p2(sK63(sK69(sK95))) | p2(sK95) | ~ r1(sK68, sK95)) | ~ spl96_1369), inference(resolution, [], [f63746, f392])).
fof(f392, plain, ! [X4, X3, X1] : (~ r1(sK69(X1), X3) | p2(X4) | ~ r1(X3, X4) | ~ p2(X3) | p2(X1) | ~ r1(sK68, X1)), inference(cnf_transformation, [], [f174])).
fof(f174, plain, (! [X1] : ((! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK69(X1), X3)) & ~ p2(sK69(X1)) & r1(X1, sK69(X1))) | p2(X1) | ~ r1(sK68, X1)) & ((! [X6] : ((r1(X6, sK71(X6)) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(sK70, X6)) & r1(sK70, sK72) & ~ p1(sK70) & r1(sK68, sK70)) | ! [X11] : ~ r1(sK68, X11) | p1(sK68)) & ((! [X13] : ((r1(X13, sK74(X13)) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(sK73, X13)) & r1(sK73, sK75) & ~ p1(sK73) & ~ p2(sK73) & r1(sK68, sK73)) | ! [X18] : ~ r1(sK68, X18) | p1(sK68) | p2(sK68)) & ((sP22(sK76) & r1(sK76, sK77) & ~ p1(sK76) & ~ p2(sK76) & ~ p3(sK76) & r1(sK68, sK76)) | ! [X21] : ~ r1(sK68, X21) | p1(sK68) | p2(sK68) | p3(sK68)) & ((sP21(sK78) & r1(sK78, sK79) & ~ p1(sK78) & ~ p2(sK78) & ~ p3(sK78) & ~ p4(sK78) & r1(sK68, sK78)) | ! [X24] : ~ r1(sK68, X24) | p1(sK68) | p2(sK68) | p3(sK68) | p4(sK68)) & ((sP19(sK80) & sP20(sK80) & ~ p1(sK80) & r1(sK68, sK80)) | ! [X26] : (! [X27] : ~ r1(X26, X27) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(sK68, X26)) | p1(sK68)) & ((sP17(sK81) & sP18(sK81) & ~ p1(sK81) & ~ p2(sK81) & r1(sK68, sK81)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(sK68, X29)) | p1(sK68) | p2(sK68)) & ((sP16(sK82) & sP15(sK82) & ~ p1(sK82) & ~ p2(sK82) & ~ p3(sK82) & r1(sK68, sK82)) | ! [X32] : (! [X33] : ~ r1(X32, X33) | p1(X32) | p2(X32) | p3(X32) | p4(X32) | ~ r1(sK68, X32)) | p1(sK68) | p2(sK68) | p3(sK68)) & ((sP13(sK83) & sP12(sK83) & ~ p1(sK83) & ~ p2(sK83) & ~ p3(sK83) & ~ p4(sK83) & r1(sK68, sK83)) | ! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(sK68, X35)) | p1(sK68) | p2(sK68) | p3(sK68) | p4(sK68)) & ((sP9(sK84) & sP10(sK84) & ~ p1(sK84) & r1(sK68, sK84)) | ! [X38] : (! [X39] : (! [X40] : ~ r1(X39, X40) | p1(X39) | p2(X39) | p3(X39) | p4(X39) | ~ r1(X38, X39)) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(sK68, X38)) | p1(sK68)) & ((! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP4(X42) | sP5(X42) | ~ r1(sK85, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(sK85, X45)) & ~ p2(sK85)) | sP2(sK85)) & r1(sK68, sK85)) | sP6(sK68)) & ! [X47] : ((p1(sK86(X47)) & (~ p1(sK87(X47)) & r1(sK86(X47), sK87(X47))) & r1(X47, sK86(X47))) | p1(X47) | ~ r1(sK68, X47)) & (~ p1(sK88) & r1(sK68, sK88)) & (sP0(sK68) | ! [X51] : ((p5(sK89(X51)) & r1(X51, sK89(X51))) | ~ r1(sK68, X51))) & ! [X53] : ((p3(sK90(X53)) & (~ p3(sK91(X53)) & r1(sK90(X53), sK91(X53))) & r1(X53, sK90(X53))) | p3(X53) | ~ r1(sK68, X53)) & (~ p3(sK92) & r1(sK68, sK92)) & ((! [X57] : ((p2(sK93(X57)) & (~ p2(sK94(X57)) & r1(sK93(X57), sK94(X57))) & r1(X57, sK93(X57))) | p2(X57) | ~ r1(sK68, X57)) & (~ p2(sK95) & r1(sK68, sK95))) | ! [X61] : (~ p5(X61) | ~ r1(sK68, X61)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK68, sK69, sK70, sK71, sK72, sK73, sK74, sK75, sK76, sK77, sK78, sK79, sK80, sK81, sK82, sK83, sK84, sK85, sK86, sK87, sK88, sK89, sK90, sK91, sK92, sK93, sK94, sK95])], [f145, f173, f172, f171, f170, f169, f168, f167, f166, f165, f164, f163, f162, f161, f160, f159, f158, f157, f156, f155, f154, f153, f152, f151, f150, f149, f148, f147, f146])).
fof(f146, plain, (? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP22(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X21] : ~ r1(X0, X21) | p1(X0) | p2(X0) | p3(X0)) & (? [X22] : (sP21(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(X0, X22)) | ! [X24] : ~ r1(X0, X24) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X25] : (sP19(X25) & sP20(X25) & ~ p1(X25) & r1(X0, X25)) | ! [X26] : (! [X27] : ~ r1(X26, X27) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | p1(X0)) & (? [X28] : (sP17(X28) & sP18(X28) & ~ p1(X28) & ~ p2(X28) & r1(X0, X28)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X0, X29)) | p1(X0) | p2(X0)) & (? [X31] : (sP16(X31) & sP15(X31) & ~ p1(X31) & ~ p2(X31) & ~ p3(X31) & r1(X0, X31)) | ! [X32] : (! [X33] : ~ r1(X32, X33) | p1(X32) | p2(X32) | p3(X32) | p4(X32) | ~ r1(X0, X32)) | p1(X0) | p2(X0) | p3(X0)) & (? [X34] : (sP13(X34) & sP12(X34) & ~ p1(X34) & ~ p2(X34) & ~ p3(X34) & ~ p4(X34) & r1(X0, X34)) | ! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(X0, X35)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X37] : (sP9(X37) & sP10(X37) & ~ p1(X37) & r1(X0, X37)) | ! [X38] : (! [X39] : (! [X40] : ~ r1(X39, X40) | p1(X39) | p2(X39) | p3(X39) | p4(X39) | ~ r1(X38, X39)) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X0, X38)) | p1(X0)) & (? [X41] : (! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP4(X42) | sP5(X42) | ~ r1(X41, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X41, X45)) & ~ p2(X41)) | sP2(X41)) & r1(X0, X41)) | sP6(X0)) & ! [X47] : (? [X48] : (p1(X48) & ? [X49] : (~ p1(X49) & r1(X48, X49)) & r1(X47, X48)) | p1(X47) | ~ r1(X0, X47)) & ? [X50] : (~ p1(X50) & r1(X0, X50)) & (sP0(X0) | ! [X51] : (? [X52] : (p5(X52) & r1(X51, X52)) | ~ r1(X0, X51))) & ! [X53] : (? [X54] : (p3(X54) & ? [X55] : (~ p3(X55) & r1(X54, X55)) & r1(X53, X54)) | p3(X53) | ~ r1(X0, X53)) & ? [X56] : (~ p3(X56) & r1(X0, X56)) & ((! [X57] : (? [X58] : (p2(X58) & ? [X59] : (~ p2(X59) & r1(X58, X59)) & r1(X57, X58)) | p2(X57) | ~ r1(X0, X57)) & ? [X60] : (~ p2(X60) & r1(X0, X60))) | ! [X61] : (~ p5(X61) | ~ r1(X0, X61)))) => (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(sK68, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(sK68, X5)) | ! [X11] : ~ r1(sK68, X11) | p1(sK68)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(sK68, X12)) | ! [X18] : ~ r1(sK68, X18) | p1(sK68) | p2(sK68)) & (? [X19] : (sP22(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(sK68, X19)) | ! [X21] : ~ r1(sK68, X21) | p1(sK68) | p2(sK68) | p3(sK68)) & (? [X22] : (sP21(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(sK68, X22)) | ! [X24] : ~ r1(sK68, X24) | p1(sK68) | p2(sK68) | p3(sK68) | p4(sK68)) & (? [X25] : (sP19(X25) & sP20(X25) & ~ p1(X25) & r1(sK68, X25)) | ! [X26] : (! [X27] : ~ r1(X26, X27) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(sK68, X26)) | p1(sK68)) & (? [X28] : (sP17(X28) & sP18(X28) & ~ p1(X28) & ~ p2(X28) & r1(sK68, X28)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(sK68, X29)) | p1(sK68) | p2(sK68)) & (? [X31] : (sP16(X31) & sP15(X31) & ~ p1(X31) & ~ p2(X31) & ~ p3(X31) & r1(sK68, X31)) | ! [X32] : (! [X33] : ~ r1(X32, X33) | p1(X32) | p2(X32) | p3(X32) | p4(X32) | ~ r1(sK68, X32)) | p1(sK68) | p2(sK68) | p3(sK68)) & (? [X34] : (sP13(X34) & sP12(X34) & ~ p1(X34) & ~ p2(X34) & ~ p3(X34) & ~ p4(X34) & r1(sK68, X34)) | ! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(sK68, X35)) | p1(sK68) | p2(sK68) | p3(sK68) | p4(sK68)) & (? [X37] : (sP9(X37) & sP10(X37) & ~ p1(X37) & r1(sK68, X37)) | ! [X38] : (! [X39] : (! [X40] : ~ r1(X39, X40) | p1(X39) | p2(X39) | p3(X39) | p4(X39) | ~ r1(X38, X39)) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(sK68, X38)) | p1(sK68)) & (? [X41] : (! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP4(X42) | sP5(X42) | ~ r1(X41, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X41, X45)) & ~ p2(X41)) | sP2(X41)) & r1(sK68, X41)) | sP6(sK68)) & ! [X47] : (? [X48] : (p1(X48) & ? [X49] : (~ p1(X49) & r1(X48, X49)) & r1(X47, X48)) | p1(X47) | ~ r1(sK68, X47)) & ? [X50] : (~ p1(X50) & r1(sK68, X50)) & (sP0(sK68) | ! [X51] : (? [X52] : (p5(X52) & r1(X51, X52)) | ~ r1(sK68, X51))) & ! [X53] : (? [X54] : (p3(X54) & ? [X55] : (~ p3(X55) & r1(X54, X55)) & r1(X53, X54)) | p3(X53) | ~ r1(sK68, X53)) & ? [X56] : (~ p3(X56) & r1(sK68, X56)) & ((! [X57] : (? [X58] : (p2(X58) & ? [X59] : (~ p2(X59) & r1(X58, X59)) & r1(X57, X58)) | p2(X57) | ~ r1(sK68, X57)) & ? [X60] : (~ p2(X60) & r1(sK68, X60))) | ! [X61] : (~ p5(X61) | ~ r1(sK68, X61))))), introduced(choice_axiom, [])).
fof(f147, plain, ! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) => (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK69(X1), X3)) & ~ p2(sK69(X1)) & r1(X1, sK69(X1)))), introduced(choice_axiom, [])).
fof(f148, plain, (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(sK68, X5)) => (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(sK70, X6)) & ? [X10] : r1(sK70, X10) & ~ p1(sK70) & r1(sK68, sK70))), introduced(choice_axiom, [])).
fof(f149, plain, ! [X6] : (? [X7] : r1(X6, X7) => r1(X6, sK71(X6))), introduced(choice_axiom, [])).
fof(f150, plain, (? [X10] : r1(sK70, X10) => r1(sK70, sK72)), introduced(choice_axiom, [])).
fof(f151, plain, (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(sK68, X12)) => (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(sK73, X13)) & ? [X17] : r1(sK73, X17) & ~ p1(sK73) & ~ p2(sK73) & r1(sK68, sK73))), introduced(choice_axiom, [])).
fof(f152, plain, ! [X13] : (? [X14] : r1(X13, X14) => r1(X13, sK74(X13))), introduced(choice_axiom, [])).
fof(f153, plain, (? [X17] : r1(sK73, X17) => r1(sK73, sK75)), introduced(choice_axiom, [])).
fof(f154, plain, (? [X19] : (sP22(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(sK68, X19)) => (sP22(sK76) & ? [X20] : r1(sK76, X20) & ~ p1(sK76) & ~ p2(sK76) & ~ p3(sK76) & r1(sK68, sK76))), introduced(choice_axiom, [])).
fof(f155, plain, (? [X20] : r1(sK76, X20) => r1(sK76, sK77)), introduced(choice_axiom, [])).
fof(f156, plain, (? [X22] : (sP21(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(sK68, X22)) => (sP21(sK78) & ? [X23] : r1(sK78, X23) & ~ p1(sK78) & ~ p2(sK78) & ~ p3(sK78) & ~ p4(sK78) & r1(sK68, sK78))), introduced(choice_axiom, [])).
fof(f157, plain, (? [X23] : r1(sK78, X23) => r1(sK78, sK79)), introduced(choice_axiom, [])).
fof(f158, plain, (? [X25] : (sP19(X25) & sP20(X25) & ~ p1(X25) & r1(sK68, X25)) => (sP19(sK80) & sP20(sK80) & ~ p1(sK80) & r1(sK68, sK80))), introduced(choice_axiom, [])).
fof(f159, plain, (? [X28] : (sP17(X28) & sP18(X28) & ~ p1(X28) & ~ p2(X28) & r1(sK68, X28)) => (sP17(sK81) & sP18(sK81) & ~ p1(sK81) & ~ p2(sK81) & r1(sK68, sK81))), introduced(choice_axiom, [])).
fof(f160, plain, (? [X31] : (sP16(X31) & sP15(X31) & ~ p1(X31) & ~ p2(X31) & ~ p3(X31) & r1(sK68, X31)) => (sP16(sK82) & sP15(sK82) & ~ p1(sK82) & ~ p2(sK82) & ~ p3(sK82) & r1(sK68, sK82))), introduced(choice_axiom, [])).
fof(f161, plain, (? [X34] : (sP13(X34) & sP12(X34) & ~ p1(X34) & ~ p2(X34) & ~ p3(X34) & ~ p4(X34) & r1(sK68, X34)) => (sP13(sK83) & sP12(sK83) & ~ p1(sK83) & ~ p2(sK83) & ~ p3(sK83) & ~ p4(sK83) & r1(sK68, sK83))), introduced(choice_axiom, [])).
fof(f162, plain, (? [X37] : (sP9(X37) & sP10(X37) & ~ p1(X37) & r1(sK68, X37)) => (sP9(sK84) & sP10(sK84) & ~ p1(sK84) & r1(sK68, sK84))), introduced(choice_axiom, [])).
fof(f163, plain, (? [X41] : (! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP4(X42) | sP5(X42) | ~ r1(X41, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X41, X45)) & ~ p2(X41)) | sP2(X41)) & r1(sK68, X41)) => (! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP4(X42) | sP5(X42) | ~ r1(sK85, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(sK85, X45)) & ~ p2(sK85)) | sP2(sK85)) & r1(sK68, sK85))), introduced(choice_axiom, [])).
fof(f164, plain, ! [X47] : (? [X48] : (p1(X48) & ? [X49] : (~ p1(X49) & r1(X48, X49)) & r1(X47, X48)) => (p1(sK86(X47)) & ? [X49] : (~ p1(X49) & r1(sK86(X47), X49)) & r1(X47, sK86(X47)))), introduced(choice_axiom, [])).
fof(f165, plain, ! [X47] : (? [X49] : (~ p1(X49) & r1(sK86(X47), X49)) => (~ p1(sK87(X47)) & r1(sK86(X47), sK87(X47)))), introduced(choice_axiom, [])).
fof(f166, plain, (? [X50] : (~ p1(X50) & r1(sK68, X50)) => (~ p1(sK88) & r1(sK68, sK88))), introduced(choice_axiom, [])).
fof(f167, plain, ! [X51] : (? [X52] : (p5(X52) & r1(X51, X52)) => (p5(sK89(X51)) & r1(X51, sK89(X51)))), introduced(choice_axiom, [])).
fof(f168, plain, ! [X53] : (? [X54] : (p3(X54) & ? [X55] : (~ p3(X55) & r1(X54, X55)) & r1(X53, X54)) => (p3(sK90(X53)) & ? [X55] : (~ p3(X55) & r1(sK90(X53), X55)) & r1(X53, sK90(X53)))), introduced(choice_axiom, [])).
fof(f169, plain, ! [X53] : (? [X55] : (~ p3(X55) & r1(sK90(X53), X55)) => (~ p3(sK91(X53)) & r1(sK90(X53), sK91(X53)))), introduced(choice_axiom, [])).
fof(f170, plain, (? [X56] : (~ p3(X56) & r1(sK68, X56)) => (~ p3(sK92) & r1(sK68, sK92))), introduced(choice_axiom, [])).
fof(f171, plain, ! [X57] : (? [X58] : (p2(X58) & ? [X59] : (~ p2(X59) & r1(X58, X59)) & r1(X57, X58)) => (p2(sK93(X57)) & ? [X59] : (~ p2(X59) & r1(sK93(X57), X59)) & r1(X57, sK93(X57)))), introduced(choice_axiom, [])).
fof(f172, plain, ! [X57] : (? [X59] : (~ p2(X59) & r1(sK93(X57), X59)) => (~ p2(sK94(X57)) & r1(sK93(X57), sK94(X57)))), introduced(choice_axiom, [])).
fof(f173, plain, (? [X60] : (~ p2(X60) & r1(sK68, X60)) => (~ p2(sK95) & r1(sK68, sK95))), introduced(choice_axiom, [])).
fof(f145, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP22(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X21] : ~ r1(X0, X21) | p1(X0) | p2(X0) | p3(X0)) & (? [X22] : (sP21(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(X0, X22)) | ! [X24] : ~ r1(X0, X24) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X25] : (sP19(X25) & sP20(X25) & ~ p1(X25) & r1(X0, X25)) | ! [X26] : (! [X27] : ~ r1(X26, X27) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | p1(X0)) & (? [X28] : (sP17(X28) & sP18(X28) & ~ p1(X28) & ~ p2(X28) & r1(X0, X28)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X0, X29)) | p1(X0) | p2(X0)) & (? [X31] : (sP16(X31) & sP15(X31) & ~ p1(X31) & ~ p2(X31) & ~ p3(X31) & r1(X0, X31)) | ! [X32] : (! [X33] : ~ r1(X32, X33) | p1(X32) | p2(X32) | p3(X32) | p4(X32) | ~ r1(X0, X32)) | p1(X0) | p2(X0) | p3(X0)) & (? [X34] : (sP13(X34) & sP12(X34) & ~ p1(X34) & ~ p2(X34) & ~ p3(X34) & ~ p4(X34) & r1(X0, X34)) | ! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(X0, X35)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X37] : (sP9(X37) & sP10(X37) & ~ p1(X37) & r1(X0, X37)) | ! [X38] : (! [X39] : (! [X40] : ~ r1(X39, X40) | p1(X39) | p2(X39) | p3(X39) | p4(X39) | ~ r1(X38, X39)) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X0, X38)) | p1(X0)) & (? [X41] : (! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP4(X42) | sP5(X42) | ~ r1(X41, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X41, X45)) & ~ p2(X41)) | sP2(X41)) & r1(X0, X41)) | sP6(X0)) & ! [X47] : (? [X48] : (p1(X48) & ? [X49] : (~ p1(X49) & r1(X48, X49)) & r1(X47, X48)) | p1(X47) | ~ r1(X0, X47)) & ? [X50] : (~ p1(X50) & r1(X0, X50)) & (sP0(X0) | ! [X51] : (? [X52] : (p5(X52) & r1(X51, X52)) | ~ r1(X0, X51))) & ! [X53] : (? [X54] : (p3(X54) & ? [X55] : (~ p3(X55) & r1(X54, X55)) & r1(X53, X54)) | p3(X53) | ~ r1(X0, X53)) & ? [X56] : (~ p3(X56) & r1(X0, X56)) & ((! [X57] : (? [X58] : (p2(X58) & ? [X59] : (~ p2(X59) & r1(X58, X59)) & r1(X57, X58)) | p2(X57) | ~ r1(X0, X57)) & ? [X60] : (~ p2(X60) & r1(X0, X60))) | ! [X61] : (~ p5(X61) | ~ r1(X0, X61)))), inference(rectify, [], [f32])).
fof(f32, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP22(X19) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (sP21(X26) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (sP19(X33) & sP20(X33) & ~ p1(X33) & r1(X0, X33)) | ! [X42] : (! [X43] : ~ r1(X42, X43) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (? [X44] : (sP17(X44) & sP18(X44) & ~ p1(X44) & ~ p2(X44) & r1(X0, X44)) | ! [X53] : (! [X54] : ~ r1(X53, X54) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (? [X55] : (sP16(X55) & sP15(X55) & ~ p1(X55) & ~ p2(X55) & ~ p3(X55) & r1(X0, X55)) | ! [X64] : (! [X65] : ~ r1(X64, X65) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (? [X66] : (sP13(X66) & sP12(X66) & ~ p1(X66) & ~ p2(X66) & ~ p3(X66) & ~ p4(X66) & r1(X0, X66)) | ! [X75] : (! [X76] : ~ r1(X75, X76) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X77] : (sP9(X77) & sP10(X77) & ~ p1(X77) & r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ~ r1(X90, X91) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (? [X92] : (! [X93] : ((! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) & ~ p2(X93)) | sP4(X93) | sP5(X93) | ~ r1(X92, X93)) & ((! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) & ~ p2(X92)) | sP2(X92)) & r1(X0, X92)) | sP6(X0)) & ! [X131] : (? [X132] : (p1(X132) & ? [X133] : (~ p1(X133) & r1(X132, X133)) & r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) & ? [X134] : (~ p1(X134) & r1(X0, X134)) & (sP0(X0) | ! [X139] : (? [X140] : (p5(X140) & r1(X139, X140)) | ~ r1(X0, X139))) & ! [X141] : (? [X142] : (p3(X142) & ? [X143] : (~ p3(X143) & r1(X142, X143)) & r1(X141, X142)) | p3(X141) | ~ r1(X0, X141)) & ? [X144] : (~ p3(X144) & r1(X0, X144)) & ((! [X145] : (? [X146] : (p2(X146) & ? [X147] : (~ p2(X147) & r1(X146, X147)) & r1(X145, X146)) | p2(X145) | ~ r1(X0, X145)) & ? [X148] : (~ p2(X148) & r1(X0, X148))) | ! [X149] : (~ p5(X149) | ~ r1(X0, X149)))), inference(definition_folding, [], [f8, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9])).
fof(f10, plain, ! [X0] : (! [X127] : (! [X128] : (? [X129] : (p2(X129) & ? [X130] : (~ p2(X130) & r1(X129, X130)) & r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)) | ~ sP1(X0)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X0] : (sP1(X0) <=> ! [X127] : (! [X128] : (? [X129] : (p2(X129) & ? [X130] : (~ p2(X130) & r1(X129, X130)) & r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f12, plain, ! [X103] : (! [X109] : (! [X110] : (? [X111] : (p2(X111) & ? [X112] : (~ p2(X112) & r1(X111, X112)) & r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)) | ~ sP3(X103)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X103] : (sP3(X103) <=> ! [X109] : (! [X110] : (? [X111] : (p2(X111) & ? [X112] : (~ p2(X112) & r1(X111, X112)) & r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f16, plain, ! [X86] : (? [X87] : (? [X88] : r1(X87, X88) & ~ p1(X87) & ~ p2(X87) & ~ p3(X87) & ~ p4(X87) & r1(X86, X87)) | ~ sP7(X86)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X86] : (sP7(X86) <=> ? [X87] : (? [X88] : r1(X87, X88) & ~ p1(X87) & ~ p2(X87) & ~ p3(X87) & ~ p4(X87) & r1(X86, X87))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f17, plain, ! [X79] : (? [X80] : (? [X81] : r1(X80, X81) & ~ p1(X80) & ~ p2(X80) & ~ p3(X80) & ~ p4(X80) & r1(X79, X80)) | ~ sP8(X79)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X79] : (sP8(X79) <=> ? [X80] : (? [X81] : r1(X80, X81) & ~ p1(X80) & ~ p2(X80) & ~ p3(X80) & ~ p4(X80) & r1(X79, X80))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f18, plain, ! [X77] : (! [X78] : ((? [X79] : (sP8(X79) & ~ p1(X79) & ~ p2(X79) & ~ p3(X79) & ~ p4(X79) & r1(X78, X79)) & ~ p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) | ~ sP9(X77)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X77] : (sP9(X77) <=> ! [X78] : ((? [X79] : (sP8(X79) & ~ p1(X79) & ~ p2(X79) & ~ p3(X79) & ~ p4(X79) & r1(X78, X79)) & ~ p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f19, plain, ! [X77] : (? [X86] : (sP7(X86) & ~ p1(X86) & ~ p2(X86) & ~ p3(X86) & ~ p4(X86) & r1(X77, X86)) | ~ sP10(X77)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X77] : (sP10(X77) <=> ? [X86] : (sP7(X86) & ~ p1(X86) & ~ p2(X86) & ~ p3(X86) & ~ p4(X86) & r1(X77, X86))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f20, plain, ! [X67] : (? [X68] : (? [X69] : r1(X68, X69) & ~ p1(X68) & ~ p2(X68) & ~ p3(X68) & ~ p4(X68) & r1(X67, X68)) | ~ sP11(X67)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X67] : (sP11(X67) <=> ? [X68] : (? [X69] : r1(X68, X69) & ~ p1(X68) & ~ p2(X68) & ~ p3(X68) & ~ p4(X68) & r1(X67, X68))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f21, plain, ! [X66] : (? [X73] : (? [X74] : r1(X73, X74) & ~ p1(X73) & ~ p2(X73) & ~ p3(X73) & ~ p4(X73) & r1(X66, X73)) | ~ sP12(X66)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X66] : (sP12(X66) <=> ? [X73] : (? [X74] : r1(X73, X74) & ~ p1(X73) & ~ p2(X73) & ~ p3(X73) & ~ p4(X73) & r1(X66, X73))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f22, plain, ! [X66] : (! [X67] : ((sP11(X67) & ~ p1(X67) & ~ p2(X67) & ~ p3(X67) & ~ p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) | ~ sP13(X66)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X66] : (sP13(X66) <=> ! [X67] : ((sP11(X67) & ~ p1(X67) & ~ p2(X67) & ~ p3(X67) & ~ p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f23, plain, ! [X56] : (? [X57] : (? [X58] : r1(X57, X58) & ~ p1(X57) & ~ p2(X57) & ~ p3(X57) & ~ p4(X57) & r1(X56, X57)) | ~ sP14(X56)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X56] : (sP14(X56) <=> ? [X57] : (? [X58] : r1(X57, X58) & ~ p1(X57) & ~ p2(X57) & ~ p3(X57) & ~ p4(X57) & r1(X56, X57))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f24, plain, ! [X55] : (? [X62] : (? [X63] : r1(X62, X63) & ~ p1(X62) & ~ p2(X62) & ~ p3(X62) & ~ p4(X62) & r1(X55, X62)) | ~ sP15(X55)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X55] : (sP15(X55) <=> ? [X62] : (? [X63] : r1(X62, X63) & ~ p1(X62) & ~ p2(X62) & ~ p3(X62) & ~ p4(X62) & r1(X55, X62))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f25, plain, ! [X55] : (! [X56] : ((sP14(X56) & ~ p1(X56) & ~ p2(X56) & ~ p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) | ~ sP16(X55)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X55] : (sP16(X55) <=> ! [X56] : ((sP14(X56) & ~ p1(X56) & ~ p2(X56) & ~ p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f26, plain, ! [X44] : (! [X45] : ((? [X46] : (? [X47] : r1(X46, X47) & ~ p1(X46) & ~ p2(X46) & ~ p3(X46) & ~ p4(X46) & r1(X45, X46)) & ~ p1(X45) & ~ p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) | ~ sP17(X44)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X44] : (sP17(X44) <=> ! [X45] : ((? [X46] : (? [X47] : r1(X46, X47) & ~ p1(X46) & ~ p2(X46) & ~ p3(X46) & ~ p4(X46) & r1(X45, X46)) & ~ p1(X45) & ~ p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f27, plain, ! [X44] : (? [X51] : (? [X52] : r1(X51, X52) & ~ p1(X51) & ~ p2(X51) & ~ p3(X51) & ~ p4(X51) & r1(X44, X51)) | ~ sP18(X44)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X44] : (sP18(X44) <=> ? [X51] : (? [X52] : r1(X51, X52) & ~ p1(X51) & ~ p2(X51) & ~ p3(X51) & ~ p4(X51) & r1(X44, X51))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f28, plain, ! [X33] : (! [X34] : ((? [X35] : (? [X36] : r1(X35, X36) & ~ p1(X35) & ~ p2(X35) & ~ p3(X35) & ~ p4(X35) & r1(X34, X35)) & ~ p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) | ~ sP19(X33)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X33] : (sP19(X33) <=> ! [X34] : ((? [X35] : (? [X36] : r1(X35, X36) & ~ p1(X35) & ~ p2(X35) & ~ p3(X35) & ~ p4(X35) & r1(X34, X35)) & ~ p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f29, plain, ! [X33] : (? [X40] : (? [X41] : r1(X40, X41) & ~ p1(X40) & ~ p2(X40) & ~ p3(X40) & ~ p4(X40) & r1(X33, X40)) | ~ sP20(X33)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X33] : (sP20(X33) <=> ? [X40] : (? [X41] : r1(X40, X41) & ~ p1(X40) & ~ p2(X40) & ~ p3(X40) & ~ p4(X40) & r1(X33, X40))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f30, plain, ! [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ~ sP21(X26)), inference(usedef, [], [e30])).
fof(e30, plain, ! [X26] : (sP21(X26) <=> ! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f31, plain, ! [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ~ sP22(X19)), inference(usedef, [], [e31])).
fof(e31, plain, ! [X19] : (sP22(X19) <=> ! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f8, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : ((? [X35] : (? [X36] : r1(X35, X36) & ~ p1(X35) & ~ p2(X35) & ~ p3(X35) & ~ p4(X35) & r1(X34, X35)) & ~ p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) & ? [X40] : (? [X41] : r1(X40, X41) & ~ p1(X40) & ~ p2(X40) & ~ p3(X40) & ~ p4(X40) & r1(X33, X40)) & ~ p1(X33) & r1(X0, X33)) | ! [X42] : (! [X43] : ~ r1(X42, X43) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (? [X44] : (! [X45] : ((? [X46] : (? [X47] : r1(X46, X47) & ~ p1(X46) & ~ p2(X46) & ~ p3(X46) & ~ p4(X46) & r1(X45, X46)) & ~ p1(X45) & ~ p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) & ? [X51] : (? [X52] : r1(X51, X52) & ~ p1(X51) & ~ p2(X51) & ~ p3(X51) & ~ p4(X51) & r1(X44, X51)) & ~ p1(X44) & ~ p2(X44) & r1(X0, X44)) | ! [X53] : (! [X54] : ~ r1(X53, X54) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (? [X55] : (! [X56] : ((? [X57] : (? [X58] : r1(X57, X58) & ~ p1(X57) & ~ p2(X57) & ~ p3(X57) & ~ p4(X57) & r1(X56, X57)) & ~ p1(X56) & ~ p2(X56) & ~ p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) & ? [X62] : (? [X63] : r1(X62, X63) & ~ p1(X62) & ~ p2(X62) & ~ p3(X62) & ~ p4(X62) & r1(X55, X62)) & ~ p1(X55) & ~ p2(X55) & ~ p3(X55) & r1(X0, X55)) | ! [X64] : (! [X65] : ~ r1(X64, X65) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (? [X66] : (! [X67] : ((? [X68] : (? [X69] : r1(X68, X69) & ~ p1(X68) & ~ p2(X68) & ~ p3(X68) & ~ p4(X68) & r1(X67, X68)) & ~ p1(X67) & ~ p2(X67) & ~ p3(X67) & ~ p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) & ? [X73] : (? [X74] : r1(X73, X74) & ~ p1(X73) & ~ p2(X73) & ~ p3(X73) & ~ p4(X73) & r1(X66, X73)) & ~ p1(X66) & ~ p2(X66) & ~ p3(X66) & ~ p4(X66) & r1(X0, X66)) | ! [X75] : (! [X76] : ~ r1(X75, X76) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X77] : (! [X78] : ((? [X79] : (? [X80] : (? [X81] : r1(X80, X81) & ~ p1(X80) & ~ p2(X80) & ~ p3(X80) & ~ p4(X80) & r1(X79, X80)) & ~ p1(X79) & ~ p2(X79) & ~ p3(X79) & ~ p4(X79) & r1(X78, X79)) & ~ p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) & ? [X86] : (? [X87] : (? [X88] : r1(X87, X88) & ~ p1(X87) & ~ p2(X87) & ~ p3(X87) & ~ p4(X87) & r1(X86, X87)) & ~ p1(X86) & ~ p2(X86) & ~ p3(X86) & ~ p4(X86) & r1(X77, X86)) & ~ p1(X77) & r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ~ r1(X90, X91) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (? [X92] : (! [X93] : ((! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) & ~ p2(X93)) | (! [X96] : (? [X97] : (p2(X97) & ? [X98] : (~ p2(X98) & r1(X97, X98)) & r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) & ? [X99] : (? [X100] : (! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) & ~ p2(X100) & r1(X99, X100)) & r1(X93, X99))) | ! [X103] : (((? [X104] : (p2(X104) & ? [X105] : (~ p2(X105) & r1(X104, X105)) & r1(X103, X104)) | p2(X103)) & (? [X106] : (! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) & ~ p2(X106) & r1(X103, X106)) | ! [X109] : (! [X110] : (? [X111] : (p2(X111) & ? [X112] : (~ p2(X112) & r1(X111, X112)) & r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)))) | ~ r1(X93, X103)) | ~ r1(X92, X93)) & ((! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) & ~ p2(X92)) | (! [X115] : (? [X116] : (p2(X116) & ? [X117] : (~ p2(X117) & r1(X116, X117)) & r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) & ? [X118] : (? [X119] : (! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) & ~ p2(X119) & r1(X118, X119)) & r1(X92, X118)))) & r1(X0, X92)) | ((? [X122] : (p2(X122) & ? [X123] : (~ p2(X123) & r1(X122, X123)) & r1(X0, X122)) | p2(X0)) & (? [X124] : (! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) & ~ p2(X124) & r1(X0, X124)) | ! [X127] : (! [X128] : (? [X129] : (p2(X129) & ? [X130] : (~ p2(X130) & r1(X129, X130)) & r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127))))) & ! [X131] : (? [X132] : (p1(X132) & ? [X133] : (~ p1(X133) & r1(X132, X133)) & r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) & ? [X134] : (~ p1(X134) & r1(X0, X134)) & ((! [X135] : (? [X136] : (p2(X136) & ? [X137] : (~ p2(X137) & r1(X136, X137)) & r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) & ? [X138] : (~ p2(X138) & r1(X0, X138))) | ! [X139] : (? [X140] : (p5(X140) & r1(X139, X140)) | ~ r1(X0, X139))) & ! [X141] : (? [X142] : (p3(X142) & ? [X143] : (~ p3(X143) & r1(X142, X143)) & r1(X141, X142)) | p3(X141) | ~ r1(X0, X141)) & ? [X144] : (~ p3(X144) & r1(X0, X144)) & ((! [X145] : (? [X146] : (p2(X146) & ? [X147] : (~ p2(X147) & r1(X146, X147)) & r1(X145, X146)) | p2(X145) | ~ r1(X0, X145)) & ? [X148] : (~ p2(X148) & r1(X0, X148))) | ! [X149] : (~ p5(X149) | ~ r1(X0, X149)))), inference(flattening, [], [f7])).
fof(f7, plain, ? [X0] : ((! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : ((? [X35] : (? [X36] : r1(X35, X36) & ~ p1(X35) & ~ p2(X35) & ~ p3(X35) & ~ p4(X35) & r1(X34, X35)) & ~ p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) & ? [X40] : (? [X41] : r1(X40, X41) & ~ p1(X40) & ~ p2(X40) & ~ p3(X40) & ~ p4(X40) & r1(X33, X40)) & ~ p1(X33) & r1(X0, X33)) | ! [X42] : (! [X43] : ~ r1(X42, X43) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (? [X44] : (! [X45] : ((? [X46] : (? [X47] : r1(X46, X47) & ~ p1(X46) & ~ p2(X46) & ~ p3(X46) & ~ p4(X46) & r1(X45, X46)) & ~ p1(X45) & ~ p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) & ? [X51] : (? [X52] : r1(X51, X52) & ~ p1(X51) & ~ p2(X51) & ~ p3(X51) & ~ p4(X51) & r1(X44, X51)) & ~ p1(X44) & ~ p2(X44) & r1(X0, X44)) | ! [X53] : (! [X54] : ~ r1(X53, X54) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (? [X55] : (! [X56] : ((? [X57] : (? [X58] : r1(X57, X58) & ~ p1(X57) & ~ p2(X57) & ~ p3(X57) & ~ p4(X57) & r1(X56, X57)) & ~ p1(X56) & ~ p2(X56) & ~ p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) & ? [X62] : (? [X63] : r1(X62, X63) & ~ p1(X62) & ~ p2(X62) & ~ p3(X62) & ~ p4(X62) & r1(X55, X62)) & ~ p1(X55) & ~ p2(X55) & ~ p3(X55) & r1(X0, X55)) | ! [X64] : (! [X65] : ~ r1(X64, X65) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (? [X66] : (! [X67] : ((? [X68] : (? [X69] : r1(X68, X69) & ~ p1(X68) & ~ p2(X68) & ~ p3(X68) & ~ p4(X68) & r1(X67, X68)) & ~ p1(X67) & ~ p2(X67) & ~ p3(X67) & ~ p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) & ? [X73] : (? [X74] : r1(X73, X74) & ~ p1(X73) & ~ p2(X73) & ~ p3(X73) & ~ p4(X73) & r1(X66, X73)) & ~ p1(X66) & ~ p2(X66) & ~ p3(X66) & ~ p4(X66) & r1(X0, X66)) | ! [X75] : (! [X76] : ~ r1(X75, X76) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X77] : (! [X78] : ((? [X79] : (? [X80] : (? [X81] : r1(X80, X81) & ~ p1(X80) & ~ p2(X80) & ~ p3(X80) & ~ p4(X80) & r1(X79, X80)) & ~ p1(X79) & ~ p2(X79) & ~ p3(X79) & ~ p4(X79) & r1(X78, X79)) & ~ p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) & ? [X86] : (? [X87] : (? [X88] : r1(X87, X88) & ~ p1(X87) & ~ p2(X87) & ~ p3(X87) & ~ p4(X87) & r1(X86, X87)) & ~ p1(X86) & ~ p2(X86) & ~ p3(X86) & ~ p4(X86) & r1(X77, X86)) & ~ p1(X77) & r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ~ r1(X90, X91) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (? [X92] : (! [X93] : (((! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) & ~ p2(X93)) | (! [X96] : (? [X97] : (p2(X97) & ? [X98] : (~ p2(X98) & r1(X97, X98)) & r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) & ? [X99] : (? [X100] : (! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) & ~ p2(X100) & r1(X99, X100)) & r1(X93, X99)))) | ! [X103] : (((? [X104] : (p2(X104) & ? [X105] : (~ p2(X105) & r1(X104, X105)) & r1(X103, X104)) | p2(X103)) & (? [X106] : (! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) & ~ p2(X106) & r1(X103, X106)) | ! [X109] : (! [X110] : (? [X111] : (p2(X111) & ? [X112] : (~ p2(X112) & r1(X111, X112)) & r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)))) | ~ r1(X93, X103)) | ~ r1(X92, X93)) & ((! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) & ~ p2(X92)) | (! [X115] : (? [X116] : (p2(X116) & ? [X117] : (~ p2(X117) & r1(X116, X117)) & r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) & ? [X118] : (? [X119] : (! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) & ~ p2(X119) & r1(X118, X119)) & r1(X92, X118)))) & r1(X0, X92)) | ((? [X122] : (p2(X122) & ? [X123] : (~ p2(X123) & r1(X122, X123)) & r1(X0, X122)) | p2(X0)) & (? [X124] : (! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) & ~ p2(X124) & r1(X0, X124)) | ! [X127] : (! [X128] : (? [X129] : (p2(X129) & ? [X130] : (~ p2(X130) & r1(X129, X130)) & r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)))))) & ! [X131] : (? [X132] : (p1(X132) & ? [X133] : (~ p1(X133) & r1(X132, X133)) & r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) & ? [X134] : (~ p1(X134) & r1(X0, X134)) & ((! [X135] : (? [X136] : (p2(X136) & ? [X137] : (~ p2(X137) & r1(X136, X137)) & r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) & ? [X138] : (~ p2(X138) & r1(X0, X138))) | ! [X139] : (? [X140] : (p5(X140) & r1(X139, X140)) | ~ r1(X0, X139))) & ! [X141] : (? [X142] : (p3(X142) & ? [X143] : (~ p3(X143) & r1(X142, X143)) & r1(X141, X142)) | p3(X141) | ~ r1(X0, X141)) & ? [X144] : (~ p3(X144) & r1(X0, X144)) & ((! [X145] : (? [X146] : (p2(X146) & ? [X147] : (~ p2(X147) & r1(X146, X147)) & r1(X145, X146)) | p2(X145) | ~ r1(X0, X145)) & ? [X148] : (~ p2(X148) & r1(X0, X148))) | ! [X149] : (~ p5(X149) | ~ r1(X0, X149)))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ~ r1(X6, X7) | p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ~ r1(X5, X10) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ~ r1(X13, X14) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ~ r1(X12, X17) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ~ r1(X20, X21) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ~ r1(X19, X24) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ~ r1(X27, X28) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ~ r1(X26, X31) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ (! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(X34, X35)) | p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) | ! [X40] : (! [X41] : ~ r1(X40, X41) | p1(X40) | p2(X40) | p3(X40) | p4(X40) | ~ r1(X33, X40)) | p1(X33) | ~ r1(X0, X33)) | ! [X42] : (! [X43] : ~ r1(X42, X43) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (~ ! [X44] : (~ ! [X45] : (~ (! [X46] : (! [X47] : ~ r1(X46, X47) | p1(X46) | p2(X46) | p3(X46) | p4(X46) | ~ r1(X45, X46)) | p1(X45) | p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) | ! [X51] : (! [X52] : ~ r1(X51, X52) | p1(X51) | p2(X51) | p3(X51) | p4(X51) | ~ r1(X44, X51)) | p1(X44) | p2(X44) | ~ r1(X0, X44)) | ! [X53] : (! [X54] : ~ r1(X53, X54) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (~ ! [X55] : (~ ! [X56] : (~ (! [X57] : (! [X58] : ~ r1(X57, X58) | p1(X57) | p2(X57) | p3(X57) | p4(X57) | ~ r1(X56, X57)) | p1(X56) | p2(X56) | p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) | ! [X62] : (! [X63] : ~ r1(X62, X63) | p1(X62) | p2(X62) | p3(X62) | p4(X62) | ~ r1(X55, X62)) | p1(X55) | p2(X55) | p3(X55) | ~ r1(X0, X55)) | ! [X64] : (! [X65] : ~ r1(X64, X65) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X66] : (~ ! [X67] : (~ (! [X68] : (! [X69] : ~ r1(X68, X69) | p1(X68) | p2(X68) | p3(X68) | p4(X68) | ~ r1(X67, X68)) | p1(X67) | p2(X67) | p3(X67) | p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) | ! [X73] : (! [X74] : ~ r1(X73, X74) | p1(X73) | p2(X73) | p3(X73) | p4(X73) | ~ r1(X66, X73)) | p1(X66) | p2(X66) | p3(X66) | p4(X66) | ~ r1(X0, X66)) | ! [X75] : (! [X76] : ~ r1(X75, X76) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X77] : (~ ! [X78] : (~ (! [X79] : (! [X80] : (! [X81] : ~ r1(X80, X81) | p1(X80) | p2(X80) | p3(X80) | p4(X80) | ~ r1(X79, X80)) | p1(X79) | p2(X79) | p3(X79) | p4(X79) | ~ r1(X78, X79)) | p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) | ! [X86] : (! [X87] : (! [X88] : ~ r1(X87, X88) | p1(X87) | p2(X87) | p3(X87) | p4(X87) | ~ r1(X86, X87)) | p1(X86) | p2(X86) | p3(X86) | p4(X86) | ~ r1(X77, X86)) | p1(X77) | ~ r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ~ r1(X90, X91) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (~ ! [X92] : (~ ! [X93] : (~ ((~ ! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | p2(X93)) & (~ ! [X96] : (~ ! [X97] : (~ p2(X97) | ! [X98] : (p2(X98) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) | ! [X99] : (! [X100] : (~ ! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) | p2(X100) | ~ r1(X99, X100)) | ~ r1(X93, X99)))) | ! [X103] : (((~ ! [X104] : (~ p2(X104) | ! [X105] : (p2(X105) | ~ r1(X104, X105)) | ~ r1(X103, X104)) | p2(X103)) & (~ ! [X106] : (~ ! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | p2(X106) | ~ r1(X103, X106)) | ! [X109] : (! [X110] : (~ ! [X111] : (~ p2(X111) | ! [X112] : (p2(X112) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)))) | ~ r1(X93, X103)) | ~ r1(X92, X93)) | ((~ ! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) | p2(X92)) & (~ ! [X115] : (~ ! [X116] : (~ p2(X116) | ! [X117] : (p2(X117) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) | ! [X118] : (! [X119] : (~ ! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) | p2(X119) | ~ r1(X118, X119)) | ~ r1(X92, X118)))) | ~ r1(X0, X92)) | ((~ ! [X122] : (~ p2(X122) | ! [X123] : (p2(X123) | ~ r1(X122, X123)) | ~ r1(X0, X122)) | p2(X0)) & (~ ! [X124] : (~ ! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | p2(X124) | ~ r1(X0, X124)) | ! [X127] : (! [X128] : (~ ! [X129] : (~ p2(X129) | ! [X130] : (p2(X130) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)))))) | ~ ! [X131] : (~ ! [X132] : (~ p1(X132) | ! [X133] : (p1(X133) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) | ! [X134] : (p1(X134) | ~ r1(X0, X134)) | ((~ ! [X135] : (~ ! [X136] : (~ p2(X136) | ! [X137] : (p2(X137) | ~ r1(X136, X137)) | ~ r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) | ! [X138] : (p2(X138) | ~ r1(X0, X138))) & ~ ! [X139] : (~ ! [X140] : (~ p5(X140) | ~ r1(X139, X140)) | ~ r1(X0, X139))) | ~ ! [X141] : (~ ! [X142] : (~ p3(X142) | ! [X143] : (p3(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | p3(X141) | ~ r1(X0, X141)) | ! [X144] : (p3(X144) | ~ r1(X0, X144)) | ((~ ! [X145] : (~ ! [X146] : (~ p2(X146) | ! [X147] : (p2(X147) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | p2(X145) | ~ r1(X0, X145)) | ! [X148] : (p2(X148) | ~ r1(X0, X148))) & ~ ! [X149] : (~ p5(X149) | ~ r1(X0, X149)))), inference(flattening, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ~ r1(X6, X7) | p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ~ r1(X5, X10) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ~ r1(X13, X14) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ~ r1(X12, X17) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ~ r1(X20, X21) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ~ r1(X19, X24) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ~ r1(X27, X28) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ~ r1(X26, X31) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ (! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(X34, X35)) | p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) | ! [X40] : (! [X41] : ~ r1(X40, X41) | p1(X40) | p2(X40) | p3(X40) | p4(X40) | ~ r1(X33, X40)) | p1(X33) | ~ r1(X0, X33)) | ! [X42] : (! [X43] : ~ r1(X42, X43) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (~ ! [X44] : (~ ! [X45] : (~ (! [X46] : (! [X47] : ~ r1(X46, X47) | p1(X46) | p2(X46) | p3(X46) | p4(X46) | ~ r1(X45, X46)) | p1(X45) | p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) | ! [X51] : (! [X52] : ~ r1(X51, X52) | p1(X51) | p2(X51) | p3(X51) | p4(X51) | ~ r1(X44, X51)) | p1(X44) | p2(X44) | ~ r1(X0, X44)) | ! [X53] : (! [X54] : ~ r1(X53, X54) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (~ ! [X55] : (~ ! [X56] : (~ (! [X57] : (! [X58] : ~ r1(X57, X58) | p1(X57) | p2(X57) | p3(X57) | p4(X57) | ~ r1(X56, X57)) | p1(X56) | p2(X56) | p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) | ! [X62] : (! [X63] : ~ r1(X62, X63) | p1(X62) | p2(X62) | p3(X62) | p4(X62) | ~ r1(X55, X62)) | p1(X55) | p2(X55) | p3(X55) | ~ r1(X0, X55)) | ! [X64] : (! [X65] : ~ r1(X64, X65) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X66] : (~ ! [X67] : (~ (! [X68] : (! [X69] : ~ r1(X68, X69) | p1(X68) | p2(X68) | p3(X68) | p4(X68) | ~ r1(X67, X68)) | p1(X67) | p2(X67) | p3(X67) | p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) | ! [X73] : (! [X74] : ~ r1(X73, X74) | p1(X73) | p2(X73) | p3(X73) | p4(X73) | ~ r1(X66, X73)) | p1(X66) | p2(X66) | p3(X66) | p4(X66) | ~ r1(X0, X66)) | ! [X75] : (! [X76] : ~ r1(X75, X76) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X77] : (~ ! [X78] : (~ (! [X79] : (! [X80] : (! [X81] : ~ r1(X80, X81) | p1(X80) | p2(X80) | p3(X80) | p4(X80) | ~ r1(X79, X80)) | p1(X79) | p2(X79) | p3(X79) | p4(X79) | ~ r1(X78, X79)) | p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) | ! [X86] : (! [X87] : (! [X88] : ~ r1(X87, X88) | p1(X87) | p2(X87) | p3(X87) | p4(X87) | ~ r1(X86, X87)) | p1(X86) | p2(X86) | p3(X86) | p4(X86) | ~ r1(X77, X86)) | p1(X77) | ~ r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ~ r1(X90, X91) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (~ ! [X92] : (~ ! [X93] : (~ ((~ ! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | p2(X93)) & (~ ! [X96] : (~ ! [X97] : (~ p2(X97) | ! [X98] : (p2(X98) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) | ! [X99] : (! [X100] : (~ ! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) | p2(X100) | ~ r1(X99, X100)) | ~ r1(X93, X99)))) | ! [X103] : (((~ ! [X104] : (~ p2(X104) | ! [X105] : (p2(X105) | ~ r1(X104, X105)) | ~ r1(X103, X104)) | p2(X103)) & (~ ! [X106] : (~ ! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | p2(X106) | ~ r1(X103, X106)) | ! [X109] : (! [X110] : (~ ! [X111] : (~ p2(X111) | ! [X112] : (p2(X112) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)))) | ~ r1(X93, X103)) | ~ r1(X92, X93)) | ((~ ! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) | p2(X92)) & (~ ! [X115] : (~ ! [X116] : (~ p2(X116) | ! [X117] : (p2(X117) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) | ! [X118] : (! [X119] : (~ ! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) | p2(X119) | ~ r1(X118, X119)) | ~ r1(X92, X118)))) | ~ r1(X0, X92)) | ((~ ! [X122] : (~ p2(X122) | ! [X123] : (p2(X123) | ~ r1(X122, X123)) | ~ r1(X0, X122)) | p2(X0)) & (~ ! [X124] : (~ ! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | p2(X124) | ~ r1(X0, X124)) | ! [X127] : (! [X128] : (~ ! [X129] : (~ p2(X129) | ! [X130] : (p2(X130) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)))))) | ~ ! [X131] : (~ ! [X132] : (~ p1(X132) | ! [X133] : (p1(X133) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) | ! [X134] : (p1(X134) | ~ r1(X0, X134)) | ((~ ! [X135] : (~ ! [X136] : (~ p2(X136) | ! [X137] : (p2(X137) | ~ r1(X136, X137)) | ~ r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) | ! [X138] : (p2(X138) | ~ r1(X0, X138))) & ~ ! [X139] : (~ ! [X140] : (~ p5(X140) | ~ r1(X139, X140)) | ~ r1(X0, X139))) | ~ ! [X141] : (~ ! [X142] : (~ p3(X142) | ! [X143] : (p3(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | p3(X141) | ~ r1(X0, X141)) | ! [X144] : (p3(X144) | ~ r1(X0, X144)) | ((~ ! [X145] : (~ ! [X146] : (~ p2(X146) | ! [X147] : (p2(X147) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | p2(X145) | ~ r1(X0, X145)) | ! [X148] : (p2(X148) | ~ r1(X0, X148))) & ~ ! [X149] : (~ p5(X149) | ~ r1(X0, X149)))), inference(true_and_false_elimination, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ($false | ~ r1(X6, X7)) | p1(X6)) | ! [X8] : (! [X9] : ($false | ~ r1(X8, X9)) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ($false | ~ r1(X5, X10)) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ($false | ~ r1(X0, X11)) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ($false | ~ r1(X13, X14)) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ($false | ~ r1(X15, X16)) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ($false | ~ r1(X12, X17)) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ($false | ~ r1(X0, X18)) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ($false | ~ r1(X20, X21)) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ($false | ~ r1(X22, X23)) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ($false | ~ r1(X19, X24)) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ($false | ~ r1(X0, X25)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ($false | ~ r1(X27, X28)) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ($false | ~ r1(X29, X30)) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ($false | ~ r1(X26, X31)) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ($false | ~ r1(X0, X32)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ (! [X35] : (! [X36] : ($false | ~ r1(X35, X36)) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(X34, X35)) | p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ($false | ~ r1(X38, X39)) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) | ! [X40] : (! [X41] : ($false | ~ r1(X40, X41)) | p1(X40) | p2(X40) | p3(X40) | p4(X40) | ~ r1(X33, X40)) | p1(X33) | ~ r1(X0, X33)) | ! [X42] : (! [X43] : ($false | ~ r1(X42, X43)) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (~ ! [X44] : (~ ! [X45] : (~ (! [X46] : (! [X47] : ($false | ~ r1(X46, X47)) | p1(X46) | p2(X46) | p3(X46) | p4(X46) | ~ r1(X45, X46)) | p1(X45) | p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ($false | ~ r1(X49, X50)) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) | ! [X51] : (! [X52] : ($false | ~ r1(X51, X52)) | p1(X51) | p2(X51) | p3(X51) | p4(X51) | ~ r1(X44, X51)) | p1(X44) | p2(X44) | ~ r1(X0, X44)) | ! [X53] : (! [X54] : ($false | ~ r1(X53, X54)) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (~ ! [X55] : (~ ! [X56] : (~ (! [X57] : (! [X58] : ($false | ~ r1(X57, X58)) | p1(X57) | p2(X57) | p3(X57) | p4(X57) | ~ r1(X56, X57)) | p1(X56) | p2(X56) | p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ($false | ~ r1(X60, X61)) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) | ! [X62] : (! [X63] : ($false | ~ r1(X62, X63)) | p1(X62) | p2(X62) | p3(X62) | p4(X62) | ~ r1(X55, X62)) | p1(X55) | p2(X55) | p3(X55) | ~ r1(X0, X55)) | ! [X64] : (! [X65] : ($false | ~ r1(X64, X65)) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X66] : (~ ! [X67] : (~ (! [X68] : (! [X69] : ($false | ~ r1(X68, X69)) | p1(X68) | p2(X68) | p3(X68) | p4(X68) | ~ r1(X67, X68)) | p1(X67) | p2(X67) | p3(X67) | p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ($false | ~ r1(X71, X72)) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) | ! [X73] : (! [X74] : ($false | ~ r1(X73, X74)) | p1(X73) | p2(X73) | p3(X73) | p4(X73) | ~ r1(X66, X73)) | p1(X66) | p2(X66) | p3(X66) | p4(X66) | ~ r1(X0, X66)) | ! [X75] : (! [X76] : ($false | ~ r1(X75, X76)) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X77] : (~ ! [X78] : (~ (! [X79] : (! [X80] : (! [X81] : ($false | ~ r1(X80, X81)) | p1(X80) | p2(X80) | p3(X80) | p4(X80) | ~ r1(X79, X80)) | p1(X79) | p2(X79) | p3(X79) | p4(X79) | ~ r1(X78, X79)) | p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ($false | ~ r1(X84, X85)) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) | ! [X86] : (! [X87] : (! [X88] : ($false | ~ r1(X87, X88)) | p1(X87) | p2(X87) | p3(X87) | p4(X87) | ~ r1(X86, X87)) | p1(X86) | p2(X86) | p3(X86) | p4(X86) | ~ r1(X77, X86)) | p1(X77) | ~ r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ($false | ~ r1(X90, X91)) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (~ ! [X92] : (~ ! [X93] : (~ ((~ ! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | p2(X93)) & (~ ! [X96] : (~ ! [X97] : (~ p2(X97) | ! [X98] : (p2(X98) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) | ! [X99] : (! [X100] : (~ ! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) | p2(X100) | ~ r1(X99, X100)) | ~ r1(X93, X99)))) | ! [X103] : (((~ ! [X104] : (~ p2(X104) | ! [X105] : (p2(X105) | ~ r1(X104, X105)) | ~ r1(X103, X104)) | p2(X103)) & (~ ! [X106] : (~ ! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | p2(X106) | ~ r1(X103, X106)) | ! [X109] : (! [X110] : (~ ! [X111] : (~ p2(X111) | ! [X112] : (p2(X112) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)))) | ~ r1(X93, X103)) | ~ r1(X92, X93)) | ((~ ! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) | p2(X92)) & (~ ! [X115] : (~ ! [X116] : (~ p2(X116) | ! [X117] : (p2(X117) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) | ! [X118] : (! [X119] : (~ ! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) | p2(X119) | ~ r1(X118, X119)) | ~ r1(X92, X118)))) | ~ r1(X0, X92)) | ((~ ! [X122] : (~ p2(X122) | ! [X123] : (p2(X123) | ~ r1(X122, X123)) | ~ r1(X0, X122)) | p2(X0)) & (~ ! [X124] : (~ ! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | p2(X124) | ~ r1(X0, X124)) | ! [X127] : (! [X128] : (~ ! [X129] : (~ p2(X129) | ! [X130] : (p2(X130) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)))))) | ~ ! [X131] : (~ ! [X132] : (~ p1(X132) | ! [X133] : (p1(X133) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) | ! [X134] : (p1(X134) | ~ r1(X0, X134)) | ((~ ! [X135] : (~ ! [X136] : (~ p2(X136) | ! [X137] : (p2(X137) | ~ r1(X136, X137)) | ~ r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) | ! [X138] : (p2(X138) | ~ r1(X0, X138))) & ~ ! [X139] : (~ ! [X140] : (~ p5(X140) | ~ r1(X139, X140)) | ~ r1(X0, X139))) | ~ ! [X141] : (~ ! [X142] : (~ p3(X142) | ! [X143] : (p3(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | p3(X141) | ~ r1(X0, X141)) | ! [X144] : (p3(X144) | ~ r1(X0, X144)) | ((~ ! [X145] : (~ ! [X146] : (~ p2(X146) | ! [X147] : (p2(X147) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | p2(X145) | ~ r1(X0, X145)) | ! [X148] : (p2(X148) | ~ r1(X0, X148))) & ~ ! [X149] : (~ p5(X149) | ~ r1(X0, X149)))), inference(rectify, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ ! [X0] : (~ p5(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ p5(X1) | ~ r1(X0, X1)))), inference(negated_conjecture, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ ! [X0] : (~ p5(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ p5(X1) | ~ r1(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL660+1.010.p', main)).
fof(f63746, plain, (r1(sK69(sK95), sK63(sK69(sK95))) | ~ spl96_1369), inference(avatar_component_clause, [], [f63744])).
fof(f97071, plain, (spl96_1362 | spl96_74 | ~ spl96_75 | ~ spl96_3938), inference(avatar_split_clause, [], [f97070, f95568, f735, f730, f63708])).
fof(f63708, plain, (spl96_1362 <=> ! [X6] : (~ r1(X6, sK95) | ~ sP1(X6))), introduced(avatar_definition, [new_symbols(naming, [spl96_1362])])).
fof(f95568, plain, (spl96_3938 <=> ! [X1, X0] : (~ r1(X0, sK69(sK95)) | ~ sP1(X1) | ~ r1(X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl96_3938])])).
fof(f97070, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK95)) | (spl96_74 | ~ spl96_75 | ~ spl96_3938)), inference(subsumption_resolution, [], [f97069, f737])).
fof(f97069, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK95) | ~ r1(sK68, sK95)) | (spl96_74 | ~ spl96_3938)), inference(subsumption_resolution, [], [f97067, f732])).
fof(f97067, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK95) | p2(sK95) | ~ r1(sK68, sK95)) | ~ spl96_3938), inference(resolution, [], [f95569, f390])).
fof(f390, plain, ! [X1] : (r1(X1, sK69(X1)) | p2(X1) | ~ r1(sK68, X1)), inference(cnf_transformation, [], [f174])).
fof(f95569, plain, (! [X0, X1] : (~ r1(X0, sK69(sK95)) | ~ sP1(X1) | ~ r1(X1, X0)) | ~ spl96_3938), inference(avatar_component_clause, [], [f95568])).
fof(f95578, plain, (spl96_74 | ~ spl96_75 | ~ spl96_3939), inference(avatar_contradiction_clause, [], [f95577])).
fof(f95577, plain, ($false | (spl96_74 | ~ spl96_75 | ~ spl96_3939)), inference(subsumption_resolution, [], [f95576, f737])).
fof(f95576, plain, (~ r1(sK68, sK95) | (spl96_74 | ~ spl96_3939)), inference(subsumption_resolution, [], [f95575, f732])).
fof(f95575, plain, (p2(sK95) | ~ r1(sK68, sK95) | ~ spl96_3939), inference(resolution, [], [f95573, f391])).
fof(f391, plain, ! [X1] : (~ p2(sK69(X1)) | p2(X1) | ~ r1(sK68, X1)), inference(cnf_transformation, [], [f174])).
fof(f95573, plain, (p2(sK69(sK95)) | ~ spl96_3939), inference(avatar_component_clause, [], [f95571])).
fof(f95571, plain, (spl96_3939 <=> p2(sK69(sK95))), introduced(avatar_definition, [new_symbols(naming, [spl96_3939])])).
fof(f95574, plain, (spl96_3938 | spl96_3939 | ~ spl96_3799), inference(avatar_split_clause, [], [f95566, f94883, f95571, f95568])).
fof(f95566, plain, (! [X0, X1] : (p2(sK69(sK95)) | ~ r1(X0, sK69(sK95)) | ~ r1(X1, X0) | ~ sP1(X1)) | ~ spl96_3799), inference(resolution, [], [f94885, f306])).
fof(f306, plain, ! [X2, X0, X1] : (~ p2(sK64(X2)) | p2(X2) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ! [X0] : (! [X1] : (! [X2] : ((p2(sK63(X2)) & (~ p2(sK64(X2)) & r1(sK63(X2), sK64(X2))) & r1(X2, sK63(X2))) | p2(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK63, sK64])], [f135, f137, f136])).
fof(f136, plain, ! [X2] : (? [X3] : (p2(X3) & ? [X4] : (~ p2(X4) & r1(X3, X4)) & r1(X2, X3)) => (p2(sK63(X2)) & ? [X4] : (~ p2(X4) & r1(sK63(X2), X4)) & r1(X2, sK63(X2)))), introduced(choice_axiom, [])).
fof(f137, plain, ! [X2] : (? [X4] : (~ p2(X4) & r1(sK63(X2), X4)) => (~ p2(sK64(X2)) & r1(sK63(X2), sK64(X2)))), introduced(choice_axiom, [])).
fof(f135, plain, ! [X0] : (! [X1] : (! [X2] : (? [X3] : (p2(X3) & ? [X4] : (~ p2(X4) & r1(X3, X4)) & r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(rectify, [], [f134])).
fof(f134, plain, ! [X0] : (! [X127] : (! [X128] : (? [X129] : (p2(X129) & ? [X130] : (~ p2(X130) & r1(X129, X130)) & r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)) | ~ sP1(X0)), inference(nnf_transformation, [], [f10])).
fof(f94885, plain, (p2(sK64(sK69(sK95))) | ~ spl96_3799), inference(avatar_component_clause, [], [f94883])).
fof(f91194, plain, (spl96_1365 | spl96_1362 | spl96_74 | ~ spl96_75), inference(avatar_split_clause, [], [f91193, f735, f730, f63708, f63722])).
fof(f91193, plain, (! [X9] : (~ r1(X9, sK95) | ~ sP1(X9) | p2(sK63(sK69(sK95)))) | (spl96_74 | ~ spl96_75)), inference(subsumption_resolution, [], [f90661, f732])).
fof(f90661, plain, (! [X9] : (~ r1(X9, sK95) | ~ sP1(X9) | p2(sK95) | p2(sK63(sK69(sK95)))) | ~ spl96_75), inference(resolution, [], [f737, f2216])).
fof(f2216, plain, ! [X30, X31] : (~ r1(sK68, X30) | ~ r1(X31, X30) | ~ sP1(X31) | p2(X30) | p2(sK63(sK69(X30)))), inference(subsumption_resolution, [], [f2182, f391])).
fof(f2182, plain, ! [X30, X31] : (p2(sK69(X30)) | p2(sK63(sK69(X30))) | ~ r1(X31, X30) | ~ sP1(X31) | p2(X30) | ~ r1(sK68, X30)), inference(resolution, [], [f307, f390])).
fof(f307, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | p2(sK63(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f138])).
fof(f91186, plain, (spl96_1369 | spl96_1362 | spl96_74 | ~ spl96_75), inference(avatar_split_clause, [], [f91185, f735, f730, f63708, f63744])).
fof(f91185, plain, (! [X17] : (~ r1(X17, sK95) | ~ sP1(X17) | r1(sK69(sK95), sK63(sK69(sK95)))) | (spl96_74 | ~ spl96_75)), inference(subsumption_resolution, [], [f90669, f732])).
fof(f90669, plain, (! [X17] : (~ r1(X17, sK95) | ~ sP1(X17) | p2(sK95) | r1(sK69(sK95), sK63(sK69(sK95)))) | ~ spl96_75), inference(resolution, [], [f737, f2506])).
fof(f2506, plain, ! [X30, X31] : (~ r1(sK68, X30) | ~ r1(X31, X30) | ~ sP1(X31) | p2(X30) | r1(sK69(X30), sK63(sK69(X30)))), inference(subsumption_resolution, [], [f2472, f391])).
fof(f2472, plain, ! [X30, X31] : (p2(sK69(X30)) | r1(sK69(X30), sK63(sK69(X30))) | ~ r1(X31, X30) | ~ sP1(X31) | p2(X30) | ~ r1(sK68, X30)), inference(resolution, [], [f304, f390])).
fof(f304, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | r1(X2, sK63(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f138])).
fof(f91174, plain, (spl96_1371 | spl96_1362 | spl96_74 | ~ spl96_75), inference(avatar_split_clause, [], [f91173, f735, f730, f63708, f63756])).
fof(f91173, plain, (! [X25] : (~ r1(X25, sK95) | ~ sP1(X25) | r1(sK63(sK69(sK95)), sK64(sK69(sK95)))) | (spl96_74 | ~ spl96_75)), inference(subsumption_resolution, [], [f90677, f732])).
fof(f90677, plain, (! [X25] : (~ r1(X25, sK95) | ~ sP1(X25) | p2(sK95) | r1(sK63(sK69(sK95)), sK64(sK69(sK95)))) | ~ spl96_75), inference(resolution, [], [f737, f2769])).
fof(f2769, plain, ! [X30, X31] : (~ r1(sK68, X30) | ~ r1(X31, X30) | ~ sP1(X31) | p2(X30) | r1(sK63(sK69(X30)), sK64(sK69(X30)))), inference(subsumption_resolution, [], [f2735, f391])).
fof(f2735, plain, ! [X30, X31] : (p2(sK69(X30)) | r1(sK63(sK69(X30)), sK64(sK69(X30))) | ~ r1(X31, X30) | ~ sP1(X31) | p2(X30) | ~ r1(sK68, X30)), inference(resolution, [], [f305, f390])).
fof(f305, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | r1(sK63(X2), sK64(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f138])).
fof(f90841, plain, (~ spl96_75 | ~ spl96_140 | ~ spl96_1362), inference(avatar_contradiction_clause, [], [f90840])).
fof(f90840, plain, ($false | (~ spl96_75 | ~ spl96_140 | ~ spl96_1362)), inference(subsumption_resolution, [], [f90678, f1168])).
fof(f1168, plain, (sP1(sK68) | ~ spl96_140), inference(avatar_component_clause, [], [f1166])).
fof(f90678, plain, (~ sP1(sK68) | (~ spl96_75 | ~ spl96_1362)), inference(resolution, [], [f737, f63709])).
fof(f63709, plain, (! [X6] : (~ r1(X6, sK95) | ~ sP1(X6)) | ~ spl96_1362), inference(avatar_component_clause, [], [f63708])).
fof(f90487, plain, (spl96_2412 | spl96_135 | ~ spl96_341 | ~ spl96_493 | ~ spl96_835 | ~ spl96_1603), inference(avatar_split_clause, [], [f87965, f71458, f52328, f42489, f26842, f1139, f85250])).
fof(f85250, plain, (spl96_2412 <=> p2(sK64(sK69(sK67(sK68))))), introduced(avatar_definition, [new_symbols(naming, [spl96_2412])])).
fof(f1139, plain, (spl96_135 <=> p2(sK67(sK68))), introduced(avatar_definition, [new_symbols(naming, [spl96_135])])).
fof(f26842, plain, (spl96_341 <=> p2(sK63(sK69(sK67(sK68))))), introduced(avatar_definition, [new_symbols(naming, [spl96_341])])).
fof(f42489, plain, (spl96_493 <=> r1(sK68, sK67(sK68))), introduced(avatar_definition, [new_symbols(naming, [spl96_493])])).
fof(f52328, plain, (spl96_835 <=> r1(sK69(sK67(sK68)), sK63(sK69(sK67(sK68))))), introduced(avatar_definition, [new_symbols(naming, [spl96_835])])).
fof(f71458, plain, (spl96_1603 <=> r1(sK63(sK69(sK67(sK68))), sK64(sK69(sK67(sK68))))), introduced(avatar_definition, [new_symbols(naming, [spl96_1603])])).
fof(f87965, plain, (p2(sK64(sK69(sK67(sK68)))) | (spl96_135 | ~ spl96_341 | ~ spl96_493 | ~ spl96_835 | ~ spl96_1603)), inference(resolution, [], [f84312, f71460])).
fof(f71460, plain, (r1(sK63(sK69(sK67(sK68))), sK64(sK69(sK67(sK68)))) | ~ spl96_1603), inference(avatar_component_clause, [], [f71458])).
fof(f84312, plain, (! [X0] : (~ r1(sK63(sK69(sK67(sK68))), X0) | p2(X0)) | (spl96_135 | ~ spl96_341 | ~ spl96_493 | ~ spl96_835)), inference(subsumption_resolution, [], [f84311, f42490])).
fof(f42490, plain, (r1(sK68, sK67(sK68)) | ~ spl96_493), inference(avatar_component_clause, [], [f42489])).
fof(f84311, plain, (! [X0] : (p2(X0) | ~ r1(sK63(sK69(sK67(sK68))), X0) | ~ r1(sK68, sK67(sK68))) | (spl96_135 | ~ spl96_341 | ~ spl96_835)), inference(subsumption_resolution, [], [f84310, f1140])).
fof(f1140, plain, (~ p2(sK67(sK68)) | spl96_135), inference(avatar_component_clause, [], [f1139])).
fof(f84310, plain, (! [X0] : (p2(X0) | ~ r1(sK63(sK69(sK67(sK68))), X0) | p2(sK67(sK68)) | ~ r1(sK68, sK67(sK68))) | (~ spl96_341 | ~ spl96_835)), inference(subsumption_resolution, [], [f84150, f26844])).
fof(f26844, plain, (p2(sK63(sK69(sK67(sK68)))) | ~ spl96_341), inference(avatar_component_clause, [], [f26842])).
fof(f84150, plain, (! [X0] : (p2(X0) | ~ r1(sK63(sK69(sK67(sK68))), X0) | ~ p2(sK63(sK69(sK67(sK68)))) | p2(sK67(sK68)) | ~ r1(sK68, sK67(sK68))) | ~ spl96_835), inference(resolution, [], [f52330, f392])).
fof(f52330, plain, (r1(sK69(sK67(sK68)), sK63(sK69(sK67(sK68)))) | ~ spl96_835), inference(avatar_component_clause, [], [f52328])).
fof(f90486, plain, (spl96_135 | ~ spl96_493 | ~ spl96_497), inference(avatar_contradiction_clause, [], [f90485])).
fof(f90485, plain, ($false | (spl96_135 | ~ spl96_493 | ~ spl96_497)), inference(subsumption_resolution, [], [f90484, f42490])).
fof(f90484, plain, (~ r1(sK68, sK67(sK68)) | (spl96_135 | ~ spl96_497)), inference(subsumption_resolution, [], [f90483, f1140])).
fof(f90483, plain, (p2(sK67(sK68)) | ~ r1(sK68, sK67(sK68)) | ~ spl96_497), inference(resolution, [], [f42506, f391])).
fof(f42506, plain, (p2(sK69(sK67(sK68))) | ~ spl96_497), inference(avatar_component_clause, [], [f42504])).
fof(f42504, plain, (spl96_497 <=> p2(sK69(sK67(sK68)))), introduced(avatar_definition, [new_symbols(naming, [spl96_497])])).
fof(f90482, plain, (spl96_342 | spl96_135 | ~ spl96_493 | ~ spl96_2699), inference(avatar_split_clause, [], [f90481, f86924, f42489, f1139, f26846])).
fof(f26846, plain, (spl96_342 <=> ! [X23] : (~ r1(X23, sK67(sK68)) | ~ sP1(X23))), introduced(avatar_definition, [new_symbols(naming, [spl96_342])])).
fof(f86924, plain, (spl96_2699 <=> ! [X1, X0] : (~ r1(X0, sK69(sK67(sK68))) | ~ sP1(X1) | ~ r1(X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl96_2699])])).
fof(f90481, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK67(sK68))) | (spl96_135 | ~ spl96_493 | ~ spl96_2699)), inference(subsumption_resolution, [], [f90480, f42490])).
fof(f90480, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK67(sK68)) | ~ r1(sK68, sK67(sK68))) | (spl96_135 | ~ spl96_2699)), inference(subsumption_resolution, [], [f90478, f1140])).
fof(f90478, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK67(sK68)) | p2(sK67(sK68)) | ~ r1(sK68, sK67(sK68))) | ~ spl96_2699), inference(resolution, [], [f86925, f390])).
fof(f86925, plain, (! [X0, X1] : (~ r1(X0, sK69(sK67(sK68))) | ~ sP1(X1) | ~ r1(X1, X0)) | ~ spl96_2699), inference(avatar_component_clause, [], [f86924])).
fof(f86926, plain, (spl96_2699 | spl96_497 | ~ spl96_2412), inference(avatar_split_clause, [], [f86922, f85250, f42504, f86924])).
fof(f86922, plain, (! [X0, X1] : (p2(sK69(sK67(sK68))) | ~ r1(X0, sK69(sK67(sK68))) | ~ r1(X1, X0) | ~ sP1(X1)) | ~ spl96_2412), inference(resolution, [], [f85252, f306])).
fof(f85252, plain, (p2(sK64(sK69(sK67(sK68)))) | ~ spl96_2412), inference(avatar_component_clause, [], [f85250])).
fof(f71463, plain, (spl96_1603 | spl96_342 | spl96_135 | ~ spl96_493), inference(avatar_split_clause, [], [f71462, f42489, f1139, f26846, f71458])).
fof(f71462, plain, (! [X1] : (~ r1(X1, sK67(sK68)) | ~ sP1(X1) | r1(sK63(sK69(sK67(sK68))), sK64(sK69(sK67(sK68))))) | (spl96_135 | ~ spl96_493)), inference(subsumption_resolution, [], [f63366, f1140])).
fof(f63366, plain, (! [X1] : (~ r1(X1, sK67(sK68)) | ~ sP1(X1) | p2(sK67(sK68)) | r1(sK63(sK69(sK67(sK68))), sK64(sK69(sK67(sK68))))) | ~ spl96_493), inference(resolution, [], [f2769, f42490])).
fof(f71092, plain, (~ spl96_59 | spl96_140 | ~ spl96_261), inference(avatar_contradiction_clause, [], [f71091])).
fof(f71091, plain, ($false | (~ spl96_59 | spl96_140 | ~ spl96_261)), inference(subsumption_resolution, [], [f71090, f671])).
fof(f671, plain, (sP6(sK68) | ~ spl96_59), inference(avatar_component_clause, [], [f669])).
fof(f71090, plain, (~ sP6(sK68) | (spl96_140 | ~ spl96_261)), inference(subsumption_resolution, [], [f71089, f1167])).
fof(f1167, plain, (~ sP1(sK68) | spl96_140), inference(avatar_component_clause, [], [f1166])).
fof(f71089, plain, (sP1(sK68) | ~ sP6(sK68) | ~ spl96_261), inference(resolution, [], [f26437, f271])).
fof(f26437, plain, (p2(sK49(sK68)) | ~ spl96_261), inference(avatar_component_clause, [], [f26435])).
fof(f26435, plain, (spl96_261 <=> p2(sK49(sK68))), introduced(avatar_definition, [new_symbols(naming, [spl96_261])])).
fof(f71025, plain, (spl96_1325 | spl96_261 | ~ spl96_1211), inference(avatar_split_clause, [], [f65059, f62030, f26435, f62935])).
fof(f62935, plain, (spl96_1325 <=> ! [X0] : (~ r1(X0, sK49(sK68)) | ~ sP0(X0))), introduced(avatar_definition, [new_symbols(naming, [spl96_1325])])).
fof(f62030, plain, (spl96_1211 <=> p2(sK66(sK49(sK68)))), introduced(avatar_definition, [new_symbols(naming, [spl96_1211])])).
fof(f65059, plain, (! [X0] : (p2(sK49(sK68)) | ~ r1(X0, sK49(sK68)) | ~ sP0(X0)) | ~ spl96_1211), inference(resolution, [], [f62032, f312])).
fof(f62032, plain, (p2(sK66(sK49(sK68))) | ~ spl96_1211), inference(avatar_component_clause, [], [f62030])).
fof(f63531, plain, (~ spl96_67 | spl96_1211 | ~ spl96_59 | spl96_140 | ~ spl96_1121 | ~ spl96_1123), inference(avatar_split_clause, [], [f63530, f54714, f54702, f1166, f669, f62030, f702])).
fof(f54702, plain, (spl96_1121 <=> r1(sK49(sK68), sK65(sK49(sK68)))), introduced(avatar_definition, [new_symbols(naming, [spl96_1121])])).
fof(f54714, plain, (spl96_1123 <=> p2(sK65(sK49(sK68)))), introduced(avatar_definition, [new_symbols(naming, [spl96_1123])])).
fof(f63530, plain, (p2(sK66(sK49(sK68))) | ~ sP0(sK68) | (~ spl96_59 | spl96_140 | ~ spl96_1121 | ~ spl96_1123)), inference(subsumption_resolution, [], [f63529, f671])).
fof(f63529, plain, (p2(sK66(sK49(sK68))) | ~ sP0(sK68) | ~ sP6(sK68) | (~ spl96_59 | spl96_140 | ~ spl96_1121 | ~ spl96_1123)), inference(subsumption_resolution, [], [f63436, f1167])).
fof(f63436, plain, (p2(sK66(sK49(sK68))) | ~ sP0(sK68) | sP1(sK68) | ~ sP6(sK68) | (~ spl96_59 | spl96_140 | ~ spl96_1121 | ~ spl96_1123)), inference(resolution, [], [f60552, f1983])).
fof(f1983, plain, ! [X11] : (r1(sK65(sK49(X11)), sK66(sK49(X11))) | ~ sP0(X11) | sP1(X11) | ~ sP6(X11)), inference(subsumption_resolution, [], [f1948, f271])).
fof(f1948, plain, ! [X11] : (p2(sK49(X11)) | r1(sK65(sK49(X11)), sK66(sK49(X11))) | ~ sP0(X11) | sP1(X11) | ~ sP6(X11)), inference(resolution, [], [f311, f270])).
fof(f60552, plain, (! [X0] : (~ r1(sK65(sK49(sK68)), X0) | p2(X0)) | (~ spl96_59 | spl96_140 | ~ spl96_1121 | ~ spl96_1123)), inference(subsumption_resolution, [], [f60551, f671])).
fof(f60551, plain, (! [X0] : (p2(X0) | ~ r1(sK65(sK49(sK68)), X0) | ~ sP6(sK68)) | (spl96_140 | ~ spl96_1121 | ~ spl96_1123)), inference(subsumption_resolution, [], [f60550, f1167])).
fof(f60550, plain, (! [X0] : (p2(X0) | ~ r1(sK65(sK49(sK68)), X0) | sP1(sK68) | ~ sP6(sK68)) | (~ spl96_1121 | ~ spl96_1123)), inference(subsumption_resolution, [], [f60419, f54716])).
fof(f54716, plain, (p2(sK65(sK49(sK68))) | ~ spl96_1123), inference(avatar_component_clause, [], [f54714])).
fof(f60419, plain, (! [X0] : (p2(X0) | ~ r1(sK65(sK49(sK68)), X0) | ~ p2(sK65(sK49(sK68))) | sP1(sK68) | ~ sP6(sK68)) | ~ spl96_1121), inference(resolution, [], [f54704, f272])).
fof(f54704, plain, (r1(sK49(sK68), sK65(sK49(sK68))) | ~ spl96_1121), inference(avatar_component_clause, [], [f54702])).
fof(f63417, plain, (~ spl96_59 | ~ spl96_67 | spl96_140 | ~ spl96_1325), inference(avatar_contradiction_clause, [], [f63416])).
fof(f63416, plain, ($false | (~ spl96_59 | ~ spl96_67 | spl96_140 | ~ spl96_1325)), inference(subsumption_resolution, [], [f63415, f671])).
fof(f63415, plain, (~ sP6(sK68) | (~ spl96_67 | spl96_140 | ~ spl96_1325)), inference(subsumption_resolution, [], [f63414, f1167])).
fof(f63414, plain, (sP1(sK68) | ~ sP6(sK68) | (~ spl96_67 | ~ spl96_1325)), inference(subsumption_resolution, [], [f63411, f704])).
fof(f63411, plain, (~ sP0(sK68) | sP1(sK68) | ~ sP6(sK68) | ~ spl96_1325), inference(resolution, [], [f62936, f270])).
fof(f62936, plain, (! [X0] : (~ r1(X0, sK49(sK68)) | ~ sP0(X0)) | ~ spl96_1325), inference(avatar_component_clause, [], [f62935])).
fof(f54717, plain, (spl96_1123 | spl96_261 | ~ spl96_67 | ~ spl96_711), inference(avatar_split_clause, [], [f54712, f51709, f702, f26435, f54714])).
fof(f51709, plain, (spl96_711 <=> r1(sK68, sK49(sK68))), introduced(avatar_definition, [new_symbols(naming, [spl96_711])])).
fof(f54712, plain, (p2(sK49(sK68)) | p2(sK65(sK49(sK68))) | (~ spl96_67 | ~ spl96_711)), inference(subsumption_resolution, [], [f54625, f704])).
fof(f54625, plain, (p2(sK49(sK68)) | p2(sK65(sK49(sK68))) | ~ sP0(sK68) | ~ spl96_711), inference(resolution, [], [f51710, f313])).
fof(f313, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK65(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f144])).
fof(f51710, plain, (r1(sK68, sK49(sK68)) | ~ spl96_711), inference(avatar_component_clause, [], [f51709])).
fof(f54705, plain, (spl96_1121 | spl96_261 | ~ spl96_67 | ~ spl96_711), inference(avatar_split_clause, [], [f54700, f51709, f702, f26435, f54702])).
fof(f54700, plain, (p2(sK49(sK68)) | r1(sK49(sK68), sK65(sK49(sK68))) | (~ spl96_67 | ~ spl96_711)), inference(subsumption_resolution, [], [f54623, f704])).
fof(f54623, plain, (p2(sK49(sK68)) | r1(sK49(sK68), sK65(sK49(sK68))) | ~ sP0(sK68) | ~ spl96_711), inference(resolution, [], [f51710, f310])).
fof(f52765, plain, (spl96_392 | spl96_64 | ~ spl96_65 | ~ spl96_67), inference(avatar_split_clause, [], [f52764, f702, f694, f689, f29775])).
fof(f52764, plain, (p2(sK65(sK85)) | (spl96_64 | ~ spl96_65 | ~ spl96_67)), inference(subsumption_resolution, [], [f52763, f704])).
fof(f52763, plain, (p2(sK65(sK85)) | ~ sP0(sK68) | (spl96_64 | ~ spl96_65)), inference(subsumption_resolution, [], [f52660, f691])).
fof(f52660, plain, (p2(sK85) | p2(sK65(sK85)) | ~ sP0(sK68) | ~ spl96_65), inference(resolution, [], [f696, f313])).
fof(f52509, plain, (spl96_341 | spl96_342 | spl96_135 | ~ spl96_493), inference(avatar_split_clause, [], [f52508, f42489, f1139, f26846, f26842])).
fof(f52508, plain, (! [X5] : (~ r1(X5, sK67(sK68)) | ~ sP1(X5) | p2(sK63(sK69(sK67(sK68))))) | (spl96_135 | ~ spl96_493)), inference(subsumption_resolution, [], [f46259, f1140])).
fof(f46259, plain, (! [X5] : (~ r1(X5, sK67(sK68)) | ~ sP1(X5) | p2(sK67(sK68)) | p2(sK63(sK69(sK67(sK68))))) | ~ spl96_493), inference(resolution, [], [f42490, f2216])).
fof(f52333, plain, (spl96_835 | spl96_342 | spl96_135 | ~ spl96_493), inference(avatar_split_clause, [], [f52332, f42489, f1139, f26846, f52328])).
fof(f52332, plain, (! [X2] : (~ r1(X2, sK67(sK68)) | ~ sP1(X2) | r1(sK69(sK67(sK68)), sK63(sK69(sK67(sK68))))) | (spl96_135 | ~ spl96_493)), inference(subsumption_resolution, [], [f52143, f1140])).
fof(f52143, plain, (! [X2] : (~ r1(X2, sK67(sK68)) | ~ sP1(X2) | p2(sK67(sK68)) | r1(sK69(sK67(sK68)), sK63(sK69(sK67(sK68))))) | ~ spl96_493), inference(resolution, [], [f2506, f42490])).
fof(f52270, plain, (~ spl96_59 | spl96_140 | spl96_711), inference(avatar_split_clause, [], [f52178, f51709, f1166, f669])).
fof(f52178, plain, (sP1(sK68) | ~ sP6(sK68) | spl96_711), inference(resolution, [], [f51711, f270])).
fof(f51711, plain, (~ r1(sK68, sK49(sK68)) | spl96_711), inference(avatar_component_clause, [], [f51709])).
fof(f46132, plain, (~ spl96_67 | spl96_493), inference(avatar_contradiction_clause, [], [f46131])).
fof(f46131, plain, ($false | (~ spl96_67 | spl96_493)), inference(subsumption_resolution, [], [f46130, f704])).
fof(f46130, plain, (~ sP0(sK68) | spl96_493), inference(resolution, [], [f42491, f308])).
fof(f308, plain, ! [X0] : (r1(X0, sK67(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f144])).
fof(f42491, plain, (~ r1(sK68, sK67(sK68)) | spl96_493), inference(avatar_component_clause, [], [f42489])).
fof(f29746, plain, (~ spl96_62 | spl96_384 | ~ spl96_385 | ~ spl96_61), inference(avatar_split_clause, [], [f29737, f677, f29743, f29739, f681])).
fof(f677, plain, (spl96_61 <=> ! [X42] : (~ p2(X42) | ~ r1(sK85, X42) | sP5(X42) | sP4(X42))), introduced(avatar_definition, [new_symbols(naming, [spl96_61])])).
fof(f29737, plain, (~ p2(sK61(sK85)) | sP4(sK61(sK85)) | ~ sP2(sK85) | ~ spl96_61), inference(subsumption_resolution, [], [f29619, f25444])).
fof(f29619, plain, (~ p2(sK61(sK85)) | sP5(sK61(sK85)) | sP4(sK61(sK85)) | ~ sP2(sK85) | ~ spl96_61), inference(resolution, [], [f678, f296])).
fof(f678, plain, (! [X42] : (~ r1(sK85, X42) | ~ p2(X42) | sP5(X42) | sP4(X42)) | ~ spl96_61), inference(avatar_component_clause, [], [f677])).
fof(f29421, plain, (~ spl96_140 | ~ spl96_67 | ~ spl96_342), inference(avatar_split_clause, [], [f29420, f26846, f702, f1166])).
fof(f29420, plain, (~ sP1(sK68) | (~ spl96_67 | ~ spl96_342)), inference(subsumption_resolution, [], [f29413, f704])).
fof(f29413, plain, (~ sP1(sK68) | ~ sP0(sK68) | ~ spl96_342), inference(resolution, [], [f26847, f308])).
fof(f26847, plain, (! [X23] : (~ r1(X23, sK67(sK68)) | ~ sP1(X23)) | ~ spl96_342), inference(avatar_component_clause, [], [f26846])).
fof(f26809, plain, (~ spl96_67 | ~ spl96_135), inference(avatar_contradiction_clause, [], [f26808])).
fof(f26808, plain, ($false | (~ spl96_67 | ~ spl96_135)), inference(subsumption_resolution, [], [f26807, f704])).
fof(f26807, plain, (~ sP0(sK68) | ~ spl96_135), inference(resolution, [], [f1141, f309])).
fof(f309, plain, ! [X0] : (~ p2(sK67(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f144])).
fof(f1141, plain, (p2(sK67(sK68)) | ~ spl96_135), inference(avatar_component_clause, [], [f1139])).
fof(f1007, plain, (~ spl96_66 | ~ spl96_68 | ~ spl96_69), inference(avatar_contradiction_clause, [], [f1006])).
fof(f1006, plain, ($false | (~ spl96_66 | ~ spl96_68 | ~ spl96_69)), inference(subsumption_resolution, [], [f1005, f978])).
fof(f978, plain, (p5(sK89(sK68)) | ~ spl96_66), inference(resolution, [], [f700, f175])).
fof(f175, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL660+1.010.p', reflexivity)).
fof(f700, plain, (! [X51] : (~ r1(sK68, X51) | p5(sK89(X51))) | ~ spl96_66), inference(avatar_component_clause, [], [f699])).
fof(f699, plain, (spl96_66 <=> ! [X51] : (p5(sK89(X51)) | ~ r1(sK68, X51))), introduced(avatar_definition, [new_symbols(naming, [spl96_66])])).
fof(f1005, plain, (~ p5(sK89(sK68)) | (~ spl96_68 | ~ spl96_69)), inference(subsumption_resolution, [], [f993, f175])).
fof(f993, plain, (~ r1(sK68, sK68) | ~ p5(sK89(sK68)) | (~ spl96_68 | ~ spl96_69)), inference(resolution, [], [f708, f712])).
fof(f712, plain, (! [X61] : (~ r1(sK68, X61) | ~ p5(X61)) | ~ spl96_69), inference(avatar_component_clause, [], [f711])).
fof(f711, plain, (spl96_69 <=> ! [X61] : (~ p5(X61) | ~ r1(sK68, X61))), introduced(avatar_definition, [new_symbols(naming, [spl96_69])])).
fof(f708, plain, (! [X51] : (r1(X51, sK89(X51)) | ~ r1(sK68, X51)) | ~ spl96_68), inference(avatar_component_clause, [], [f707])).
fof(f707, plain, (spl96_68 <=> ! [X51] : (r1(X51, sK89(X51)) | ~ r1(sK68, X51))), introduced(avatar_definition, [new_symbols(naming, [spl96_68])])).
fof(f738, plain, (spl96_69 | spl96_75), inference(avatar_split_clause, [], [f314, f735, f711])).
fof(f314, plain, ! [X61] : (r1(sK68, sK95) | ~ p5(X61) | ~ r1(sK68, X61)), inference(cnf_transformation, [], [f174])).
fof(f733, plain, (spl96_69 | ~ spl96_74), inference(avatar_split_clause, [], [f315, f730, f711])).
fof(f315, plain, ! [X61] : (~ p2(sK95) | ~ p5(X61) | ~ r1(sK68, X61)), inference(cnf_transformation, [], [f174])).
fof(f728, plain, (spl96_69 | spl96_73), inference(avatar_split_clause, [], [f316, f726, f711])).
fof(f316, plain, ! [X61, X57] : (r1(X57, sK93(X57)) | p2(X57) | ~ r1(sK68, X57) | ~ p5(X61) | ~ r1(sK68, X61)), inference(cnf_transformation, [], [f174])).
fof(f724, plain, (spl96_69 | spl96_72), inference(avatar_split_clause, [], [f317, f722, f711])).
fof(f317, plain, ! [X61, X57] : (r1(sK93(X57), sK94(X57)) | p2(X57) | ~ r1(sK68, X57) | ~ p5(X61) | ~ r1(sK68, X61)), inference(cnf_transformation, [], [f174])).
fof(f720, plain, (spl96_69 | spl96_71), inference(avatar_split_clause, [], [f318, f718, f711])).
fof(f318, plain, ! [X61, X57] : (~ p2(sK94(X57)) | p2(X57) | ~ r1(sK68, X57) | ~ p5(X61) | ~ r1(sK68, X61)), inference(cnf_transformation, [], [f174])).
fof(f716, plain, (spl96_69 | spl96_70), inference(avatar_split_clause, [], [f319, f714, f711])).
fof(f319, plain, ! [X61, X57] : (p2(sK93(X57)) | p2(X57) | ~ r1(sK68, X57) | ~ p5(X61) | ~ r1(sK68, X61)), inference(cnf_transformation, [], [f174])).
fof(f709, plain, (spl96_68 | spl96_67), inference(avatar_split_clause, [], [f326, f702, f707])).
fof(f326, plain, ! [X51] : (sP0(sK68) | r1(X51, sK89(X51)) | ~ r1(sK68, X51)), inference(cnf_transformation, [], [f174])).
fof(f705, plain, (spl96_66 | spl96_67), inference(avatar_split_clause, [], [f327, f702, f699])).
fof(f327, plain, ! [X51] : (sP0(sK68) | p5(sK89(X51)) | ~ r1(sK68, X51)), inference(cnf_transformation, [], [f174])).
fof(f697, plain, (spl96_59 | spl96_65), inference(avatar_split_clause, [], [f334, f694, f669])).
fof(f334, plain, (r1(sK68, sK85) | sP6(sK68)), inference(cnf_transformation, [], [f174])).
fof(f692, plain, (spl96_59 | spl96_62 | ~ spl96_64), inference(avatar_split_clause, [], [f335, f689, f681, f669])).
fof(f335, plain, (~ p2(sK85) | sP2(sK85) | sP6(sK68)), inference(cnf_transformation, [], [f174])).
fof(f687, plain, (spl96_59 | spl96_62 | spl96_63), inference(avatar_split_clause, [], [f336, f685, f681, f669])).
fof(f336, plain, ! [X45, X46] : (~ p2(X45) | p2(X46) | ~ r1(X45, X46) | ~ r1(sK85, X45) | sP2(sK85) | sP6(sK68)), inference(cnf_transformation, [], [f174])).
fof(f679, plain, (spl96_59 | spl96_61), inference(avatar_split_clause, [], [f337, f677, f669])).
fof(f337, plain, ! [X42] : (~ p2(X42) | sP4(X42) | sP5(X42) | ~ r1(sK85, X42) | sP6(sK68)), inference(cnf_transformation, [], [f174])).
fof(f675, plain, (spl96_59 | spl96_60), inference(avatar_split_clause, [], [f338, f673, f669])).
fof(f338, plain, ! [X43, X44, X42] : (~ p2(X43) | p2(X44) | ~ r1(X43, X44) | ~ r1(X42, X43) | sP4(X42) | sP5(X42) | ~ r1(sK85, X42) | sP6(sK68)), inference(cnf_transformation, [], [f174])).