fof(f49617, plain, $false, inference(avatar_sat_refutation, [], [f347, f351, f359, f364, f369, f377, f381, f388, f392, f396, f400, f405, f410, f517, f6879, f6885, f6891, f7138, f7178, f11272, f11310, f13683, f13948, f15963, f16341, f16918, f16940, f16952, f20840, f21092, f24417, f29078, f29082, f29083, f30502, f30506, f32137, f32151, f32316, f32910, f32913, f33031, f33161, f33328, f37873, f37889, f37895, f37919, f40113, f40335, f40369, f49616])).
fof(f49616, plain, (spl55_290 | spl55_80 | ~ spl55_399 | ~ spl55_1692), inference(avatar_split_clause, [], [f49615, f30496, f13439, f622, f7176])).
fof(f7176, plain, (spl55_290 <=> ! [X14] : (~ r1(X14, sK31(sK32)) | ~ sP1(X14))), introduced(avatar_definition, [new_symbols(naming, [spl55_290])])).
fof(f622, plain, (spl55_80 <=> p2(sK31(sK32))), introduced(avatar_definition, [new_symbols(naming, [spl55_80])])).
fof(f13439, plain, (spl55_399 <=> r1(sK32, sK31(sK32))), introduced(avatar_definition, [new_symbols(naming, [spl55_399])])).
fof(f30496, plain, (spl55_1692 <=> ! [X1, X0] : (~ r1(X0, sK33(sK31(sK32))) | ~ sP1(X1) | ~ r1(X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl55_1692])])).
fof(f49615, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK31(sK32))) | (spl55_80 | ~ spl55_399 | ~ spl55_1692)), inference(subsumption_resolution, [], [f49614, f13440])).
fof(f13440, plain, (r1(sK32, sK31(sK32)) | ~ spl55_399), inference(avatar_component_clause, [], [f13439])).
fof(f49614, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK31(sK32)) | ~ r1(sK32, sK31(sK32))) | (spl55_80 | ~ spl55_1692)), inference(subsumption_resolution, [], [f49612, f623])).
fof(f623, plain, (~ p2(sK31(sK32)) | spl55_80), inference(avatar_component_clause, [], [f622])).
fof(f49612, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK31(sK32)) | p2(sK31(sK32)) | ~ r1(sK32, sK31(sK32))) | ~ spl55_1692), inference(resolution, [], [f30497, f198])).
fof(f198, plain, ! [X1] : (r1(X1, sK33(X1)) | p2(X1) | ~ r1(sK32, X1)), inference(cnf_transformation, [], [f93])).
fof(f93, plain, (! [X1] : ((! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK33(X1), X3)) & ~ p2(sK33(X1)) & r1(X1, sK33(X1))) | p2(X1) | ~ r1(sK32, X1)) & ((! [X6] : ((r1(X6, sK35(X6)) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(sK34, X6)) & r1(sK34, sK36) & ~ p1(sK34) & r1(sK32, sK34)) | ! [X11] : ~ r1(sK32, X11) | p1(sK32)) & ((! [X13] : ((r1(X13, sK38(X13)) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(sK37, X13)) & r1(sK37, sK39) & ~ p1(sK37) & ~ p2(sK37) & r1(sK32, sK37)) | ! [X18] : ~ r1(sK32, X18) | p1(sK32) | p2(sK32)) & ((sP8(sK40) & r1(sK40, sK41) & ~ p1(sK40) & ~ p2(sK40) & ~ p3(sK40) & r1(sK32, sK40)) | ! [X21] : ~ r1(sK32, X21) | p1(sK32) | p2(sK32) | p3(sK32)) & ((sP7(sK42) & r1(sK42, sK43) & ~ p1(sK42) & ~ p2(sK42) & ~ p3(sK42) & ~ p4(sK42) & r1(sK32, sK42)) | ! [X24] : ~ r1(sK32, X24) | p1(sK32) | p2(sK32) | p3(sK32) | p4(sK32)) & ((! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP4(X26) | sP5(X26) | ~ r1(sK44, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(sK44, X29)) & ~ p2(sK44)) | sP2(sK44)) & r1(sK32, sK44)) | sP6(sK32)) & ! [X31] : ((p1(sK45(X31)) & (~ p1(sK46(X31)) & r1(sK45(X31), sK46(X31))) & r1(X31, sK45(X31))) | p1(X31) | ~ r1(sK32, X31)) & (~ p1(sK47) & r1(sK32, sK47)) & (sP0(sK32) | ! [X35] : ((p5(sK48(X35)) & r1(X35, sK48(X35))) | ~ r1(sK32, X35))) & ! [X37] : ((p3(sK49(X37)) & (~ p3(sK50(X37)) & r1(sK49(X37), sK50(X37))) & r1(X37, sK49(X37))) | p3(X37) | ~ r1(sK32, X37)) & (~ p3(sK51) & r1(sK32, sK51)) & ((! [X41] : ((p2(sK52(X41)) & (~ p2(sK53(X41)) & r1(sK52(X41), sK53(X41))) & r1(X41, sK52(X41))) | p2(X41) | ~ r1(sK32, X41)) & (~ p2(sK54) & r1(sK32, sK54))) | ! [X45] : (~ p5(X45) | ~ r1(sK32, X45)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK32, sK33, sK34, sK35, sK36, sK37, sK38, sK39, sK40, sK41, sK42, sK43, sK44, sK45, sK46, sK47, sK48, sK49, sK50, sK51, sK52, sK53, sK54])], [f69, f92, f91, f90, f89, f88, f87, f86, f85, f84, f83, f82, f81, f80, f79, f78, f77, f76, f75, f74, f73, f72, f71, f70])).
fof(f70, plain, (? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP8(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X21] : ~ r1(X0, X21) | p1(X0) | p2(X0) | p3(X0)) & (? [X22] : (sP7(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(X0, X22)) | ! [X24] : ~ r1(X0, X24) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP4(X26) | sP5(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP2(X25)) & r1(X0, X25)) | sP6(X0)) & ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) | p1(X31) | ~ r1(X0, X31)) & ? [X34] : (~ p1(X34) & r1(X0, X34)) & (sP0(X0) | ! [X35] : (? [X36] : (p5(X36) & r1(X35, X36)) | ~ r1(X0, X35))) & ! [X37] : (? [X38] : (p3(X38) & ? [X39] : (~ p3(X39) & r1(X38, X39)) & r1(X37, X38)) | p3(X37) | ~ r1(X0, X37)) & ? [X40] : (~ p3(X40) & r1(X0, X40)) & ((! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X0, X41)) & ? [X44] : (~ p2(X44) & r1(X0, X44))) | ! [X45] : (~ p5(X45) | ~ r1(X0, X45)))) => (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(sK32, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(sK32, X5)) | ! [X11] : ~ r1(sK32, X11) | p1(sK32)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(sK32, X12)) | ! [X18] : ~ r1(sK32, X18) | p1(sK32) | p2(sK32)) & (? [X19] : (sP8(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(sK32, X19)) | ! [X21] : ~ r1(sK32, X21) | p1(sK32) | p2(sK32) | p3(sK32)) & (? [X22] : (sP7(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(sK32, X22)) | ! [X24] : ~ r1(sK32, X24) | p1(sK32) | p2(sK32) | p3(sK32) | p4(sK32)) & (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP4(X26) | sP5(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP2(X25)) & r1(sK32, X25)) | sP6(sK32)) & ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) | p1(X31) | ~ r1(sK32, X31)) & ? [X34] : (~ p1(X34) & r1(sK32, X34)) & (sP0(sK32) | ! [X35] : (? [X36] : (p5(X36) & r1(X35, X36)) | ~ r1(sK32, X35))) & ! [X37] : (? [X38] : (p3(X38) & ? [X39] : (~ p3(X39) & r1(X38, X39)) & r1(X37, X38)) | p3(X37) | ~ r1(sK32, X37)) & ? [X40] : (~ p3(X40) & r1(sK32, X40)) & ((! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(sK32, X41)) & ? [X44] : (~ p2(X44) & r1(sK32, X44))) | ! [X45] : (~ p5(X45) | ~ r1(sK32, X45))))), introduced(choice_axiom, [])).
fof(f71, plain, ! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) => (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK33(X1), X3)) & ~ p2(sK33(X1)) & r1(X1, sK33(X1)))), introduced(choice_axiom, [])).
fof(f72, plain, (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(sK32, X5)) => (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(sK34, X6)) & ? [X10] : r1(sK34, X10) & ~ p1(sK34) & r1(sK32, sK34))), introduced(choice_axiom, [])).
fof(f73, plain, ! [X6] : (? [X7] : r1(X6, X7) => r1(X6, sK35(X6))), introduced(choice_axiom, [])).
fof(f74, plain, (? [X10] : r1(sK34, X10) => r1(sK34, sK36)), introduced(choice_axiom, [])).
fof(f75, plain, (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(sK32, X12)) => (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(sK37, X13)) & ? [X17] : r1(sK37, X17) & ~ p1(sK37) & ~ p2(sK37) & r1(sK32, sK37))), introduced(choice_axiom, [])).
fof(f76, plain, ! [X13] : (? [X14] : r1(X13, X14) => r1(X13, sK38(X13))), introduced(choice_axiom, [])).
fof(f77, plain, (? [X17] : r1(sK37, X17) => r1(sK37, sK39)), introduced(choice_axiom, [])).
fof(f78, plain, (? [X19] : (sP8(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(sK32, X19)) => (sP8(sK40) & ? [X20] : r1(sK40, X20) & ~ p1(sK40) & ~ p2(sK40) & ~ p3(sK40) & r1(sK32, sK40))), introduced(choice_axiom, [])).
fof(f79, plain, (? [X20] : r1(sK40, X20) => r1(sK40, sK41)), introduced(choice_axiom, [])).
fof(f80, plain, (? [X22] : (sP7(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(sK32, X22)) => (sP7(sK42) & ? [X23] : r1(sK42, X23) & ~ p1(sK42) & ~ p2(sK42) & ~ p3(sK42) & ~ p4(sK42) & r1(sK32, sK42))), introduced(choice_axiom, [])).
fof(f81, plain, (? [X23] : r1(sK42, X23) => r1(sK42, sK43)), introduced(choice_axiom, [])).
fof(f82, plain, (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP4(X26) | sP5(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP2(X25)) & r1(sK32, X25)) => (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP4(X26) | sP5(X26) | ~ r1(sK44, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(sK44, X29)) & ~ p2(sK44)) | sP2(sK44)) & r1(sK32, sK44))), introduced(choice_axiom, [])).
fof(f83, plain, ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) => (p1(sK45(X31)) & ? [X33] : (~ p1(X33) & r1(sK45(X31), X33)) & r1(X31, sK45(X31)))), introduced(choice_axiom, [])).
fof(f84, plain, ! [X31] : (? [X33] : (~ p1(X33) & r1(sK45(X31), X33)) => (~ p1(sK46(X31)) & r1(sK45(X31), sK46(X31)))), introduced(choice_axiom, [])).
fof(f85, plain, (? [X34] : (~ p1(X34) & r1(sK32, X34)) => (~ p1(sK47) & r1(sK32, sK47))), introduced(choice_axiom, [])).
fof(f86, plain, ! [X35] : (? [X36] : (p5(X36) & r1(X35, X36)) => (p5(sK48(X35)) & r1(X35, sK48(X35)))), introduced(choice_axiom, [])).
fof(f87, plain, ! [X37] : (? [X38] : (p3(X38) & ? [X39] : (~ p3(X39) & r1(X38, X39)) & r1(X37, X38)) => (p3(sK49(X37)) & ? [X39] : (~ p3(X39) & r1(sK49(X37), X39)) & r1(X37, sK49(X37)))), introduced(choice_axiom, [])).
fof(f88, plain, ! [X37] : (? [X39] : (~ p3(X39) & r1(sK49(X37), X39)) => (~ p3(sK50(X37)) & r1(sK49(X37), sK50(X37)))), introduced(choice_axiom, [])).
fof(f89, plain, (? [X40] : (~ p3(X40) & r1(sK32, X40)) => (~ p3(sK51) & r1(sK32, sK51))), introduced(choice_axiom, [])).
fof(f90, plain, ! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) => (p2(sK52(X41)) & ? [X43] : (~ p2(X43) & r1(sK52(X41), X43)) & r1(X41, sK52(X41)))), introduced(choice_axiom, [])).
fof(f91, plain, ! [X41] : (? [X43] : (~ p2(X43) & r1(sK52(X41), X43)) => (~ p2(sK53(X41)) & r1(sK52(X41), sK53(X41)))), introduced(choice_axiom, [])).
fof(f92, plain, (? [X44] : (~ p2(X44) & r1(sK32, X44)) => (~ p2(sK54) & r1(sK32, sK54))), introduced(choice_axiom, [])).
fof(f69, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP8(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X21] : ~ r1(X0, X21) | p1(X0) | p2(X0) | p3(X0)) & (? [X22] : (sP7(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(X0, X22)) | ! [X24] : ~ r1(X0, X24) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP4(X26) | sP5(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP2(X25)) & r1(X0, X25)) | sP6(X0)) & ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) | p1(X31) | ~ r1(X0, X31)) & ? [X34] : (~ p1(X34) & r1(X0, X34)) & (sP0(X0) | ! [X35] : (? [X36] : (p5(X36) & r1(X35, X36)) | ~ r1(X0, X35))) & ! [X37] : (? [X38] : (p3(X38) & ? [X39] : (~ p3(X39) & r1(X38, X39)) & r1(X37, X38)) | p3(X37) | ~ r1(X0, X37)) & ? [X40] : (~ p3(X40) & r1(X0, X40)) & ((! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X0, X41)) & ? [X44] : (~ p2(X44) & r1(X0, X44))) | ! [X45] : (~ p5(X45) | ~ r1(X0, X45)))), inference(rectify, [], [f18])).
fof(f18, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP8(X19) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (sP7(X26) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : ((! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p2(X34)) | sP4(X34) | sP5(X34) | ~ r1(X33, X34)) & ((! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) & ~ p2(X33)) | sP2(X33)) & r1(X0, X33)) | sP6(X0)) & ! [X72] : (? [X73] : (p1(X73) & ? [X74] : (~ p1(X74) & r1(X73, X74)) & r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) & ? [X75] : (~ p1(X75) & r1(X0, X75)) & (sP0(X0) | ! [X80] : (? [X81] : (p5(X81) & r1(X80, X81)) | ~ r1(X0, X80))) & ! [X82] : (? [X83] : (p3(X83) & ? [X84] : (~ p3(X84) & r1(X83, X84)) & r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) & ? [X85] : (~ p3(X85) & r1(X0, X85)) & ((! [X86] : (? [X87] : (p2(X87) & ? [X88] : (~ p2(X88) & r1(X87, X88)) & r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) & ? [X89] : (~ p2(X89) & r1(X0, X89))) | ! [X90] : (~ p5(X90) | ~ r1(X0, X90)))), inference(definition_folding, [], [f8, e17, e16, e15, e14, e13, e12, e11, e10, e9])).
fof(f9, plain, ! [X0] : ((! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79))) | ~ sP0(X0)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X0] : (sP0(X0) <=> (! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f10, plain, ! [X0] : (! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)) | ~ sP1(X0)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X0] : (sP1(X0) <=> ! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f11, plain, ! [X33] : ((! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59))) | ~ sP2(X33)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X33] : (sP2(X33) <=> (! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f12, plain, ! [X44] : (! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)) | ~ sP3(X44)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X44] : (sP3(X44) <=> ! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f13, plain, ! [X34] : ((! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40))) | ~ sP4(X34)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X34] : (sP4(X34) <=> (! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f14, plain, ! [X34] : (! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | sP3(X44))) | ~ r1(X34, X44)) | ~ sP5(X34)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X34] : (sP5(X34) <=> ! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | sP3(X44))) | ~ r1(X34, X44))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f15, plain, ! [X0] : (((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | sP1(X0))) | ~ sP6(X0)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X0] : (sP6(X0) <=> ((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | sP1(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f16, plain, ! [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ~ sP7(X26)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X26] : (sP7(X26) <=> ! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f17, plain, ! [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ~ sP8(X19)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X19] : (sP8(X19) <=> ! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f8, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : ((! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p2(X34)) | (! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40))) | ! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | ! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) & ((! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) & ~ p2(X33)) | (! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59)))) & r1(X0, X33)) | ((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | ! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68))))) & ! [X72] : (? [X73] : (p1(X73) & ? [X74] : (~ p1(X74) & r1(X73, X74)) & r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) & ? [X75] : (~ p1(X75) & r1(X0, X75)) & ((! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79))) | ! [X80] : (? [X81] : (p5(X81) & r1(X80, X81)) | ~ r1(X0, X80))) & ! [X82] : (? [X83] : (p3(X83) & ? [X84] : (~ p3(X84) & r1(X83, X84)) & r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) & ? [X85] : (~ p3(X85) & r1(X0, X85)) & ((! [X86] : (? [X87] : (p2(X87) & ? [X88] : (~ p2(X88) & r1(X87, X88)) & r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) & ? [X89] : (~ p2(X89) & r1(X0, X89))) | ! [X90] : (~ p5(X90) | ~ r1(X0, X90)))), inference(flattening, [], [f7])).
fof(f7, plain, ? [X0] : ((! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : (((! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p2(X34)) | (! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40)))) | ! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | ! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) & ((! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) & ~ p2(X33)) | (! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59)))) & r1(X0, X33)) | ((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | ! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) & ! [X72] : (? [X73] : (p1(X73) & ? [X74] : (~ p1(X74) & r1(X73, X74)) & r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) & ? [X75] : (~ p1(X75) & r1(X0, X75)) & ((! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79))) | ! [X80] : (? [X81] : (p5(X81) & r1(X80, X81)) | ~ r1(X0, X80))) & ! [X82] : (? [X83] : (p3(X83) & ? [X84] : (~ p3(X84) & r1(X83, X84)) & r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) & ? [X85] : (~ p3(X85) & r1(X0, X85)) & ((! [X86] : (? [X87] : (p2(X87) & ? [X88] : (~ p2(X88) & r1(X87, X88)) & r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) & ? [X89] : (~ p2(X89) & r1(X0, X89))) | ! [X90] : (~ p5(X90) | ~ r1(X0, X90)))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ~ r1(X6, X7) | p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ~ r1(X5, X10) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ~ r1(X13, X14) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ~ r1(X12, X17) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ~ r1(X20, X21) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ~ r1(X19, X24) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ~ r1(X27, X28) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ~ r1(X26, X31) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p2(X34)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X34, X40)))) | ! [X44] : (((~ ! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p2(X44)) & (~ ! [X47] : (~ ! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | p2(X47) | ~ r1(X44, X47)) | ! [X50] : (! [X51] : (~ ! [X52] : (~ p2(X52) | ! [X53] : (p2(X53) | ~ r1(X52, X53)) | ~ r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) | ((~ ! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) | p2(X33)) & (~ ! [X56] : (~ ! [X57] : (~ p2(X57) | ! [X58] : (p2(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) | ! [X59] : (! [X60] : (~ ! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | p2(X60) | ~ r1(X59, X60)) | ~ r1(X33, X59)))) | ~ r1(X0, X33)) | ((~ ! [X63] : (~ p2(X63) | ! [X64] : (p2(X64) | ~ r1(X63, X64)) | ~ r1(X0, X63)) | p2(X0)) & (~ ! [X65] : (~ ! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) | p2(X65) | ~ r1(X0, X65)) | ! [X68] : (! [X69] : (~ ! [X70] : (~ p2(X70) | ! [X71] : (p2(X71) | ~ r1(X70, X71)) | ~ r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) | ~ ! [X72] : (~ ! [X73] : (~ p1(X73) | ! [X74] : (p1(X74) | ~ r1(X73, X74)) | ~ r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) | ! [X75] : (p1(X75) | ~ r1(X0, X75)) | ((~ ! [X76] : (~ ! [X77] : (~ p2(X77) | ! [X78] : (p2(X78) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) | ! [X79] : (p2(X79) | ~ r1(X0, X79))) & ~ ! [X80] : (~ ! [X81] : (~ p5(X81) | ~ r1(X80, X81)) | ~ r1(X0, X80))) | ~ ! [X82] : (~ ! [X83] : (~ p3(X83) | ! [X84] : (p3(X84) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) | ! [X85] : (p3(X85) | ~ r1(X0, X85)) | ((~ ! [X86] : (~ ! [X87] : (~ p2(X87) | ! [X88] : (p2(X88) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) | ! [X89] : (p2(X89) | ~ r1(X0, X89))) & ~ ! [X90] : (~ p5(X90) | ~ r1(X0, X90)))), inference(flattening, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ~ r1(X6, X7) | p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ~ r1(X5, X10) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ~ r1(X13, X14) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ~ r1(X12, X17) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ~ r1(X20, X21) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ~ r1(X19, X24) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ~ r1(X27, X28) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ~ r1(X26, X31) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p2(X34)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X34, X40)))) | ! [X44] : (((~ ! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p2(X44)) & (~ ! [X47] : (~ ! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | p2(X47) | ~ r1(X44, X47)) | ! [X50] : (! [X51] : (~ ! [X52] : (~ p2(X52) | ! [X53] : (p2(X53) | ~ r1(X52, X53)) | ~ r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) | ((~ ! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) | p2(X33)) & (~ ! [X56] : (~ ! [X57] : (~ p2(X57) | ! [X58] : (p2(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) | ! [X59] : (! [X60] : (~ ! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | p2(X60) | ~ r1(X59, X60)) | ~ r1(X33, X59)))) | ~ r1(X0, X33)) | ((~ ! [X63] : (~ p2(X63) | ! [X64] : (p2(X64) | ~ r1(X63, X64)) | ~ r1(X0, X63)) | p2(X0)) & (~ ! [X65] : (~ ! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) | p2(X65) | ~ r1(X0, X65)) | ! [X68] : (! [X69] : (~ ! [X70] : (~ p2(X70) | ! [X71] : (p2(X71) | ~ r1(X70, X71)) | ~ r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) | ~ ! [X72] : (~ ! [X73] : (~ p1(X73) | ! [X74] : (p1(X74) | ~ r1(X73, X74)) | ~ r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) | ! [X75] : (p1(X75) | ~ r1(X0, X75)) | ((~ ! [X76] : (~ ! [X77] : (~ p2(X77) | ! [X78] : (p2(X78) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) | ! [X79] : (p2(X79) | ~ r1(X0, X79))) & ~ ! [X80] : (~ ! [X81] : (~ p5(X81) | ~ r1(X80, X81)) | ~ r1(X0, X80))) | ~ ! [X82] : (~ ! [X83] : (~ p3(X83) | ! [X84] : (p3(X84) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) | ! [X85] : (p3(X85) | ~ r1(X0, X85)) | ((~ ! [X86] : (~ ! [X87] : (~ p2(X87) | ! [X88] : (p2(X88) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) | ! [X89] : (p2(X89) | ~ r1(X0, X89))) & ~ ! [X90] : (~ p5(X90) | ~ r1(X0, X90)))), inference(true_and_false_elimination, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ($false | ~ r1(X6, X7)) | p1(X6)) | ! [X8] : (! [X9] : ($false | ~ r1(X8, X9)) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ($false | ~ r1(X5, X10)) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ($false | ~ r1(X0, X11)) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ($false | ~ r1(X13, X14)) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ($false | ~ r1(X15, X16)) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ($false | ~ r1(X12, X17)) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ($false | ~ r1(X0, X18)) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ($false | ~ r1(X20, X21)) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ($false | ~ r1(X22, X23)) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ($false | ~ r1(X19, X24)) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ($false | ~ r1(X0, X25)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ($false | ~ r1(X27, X28)) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ($false | ~ r1(X29, X30)) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ($false | ~ r1(X26, X31)) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ($false | ~ r1(X0, X32)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p2(X34)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X34, X40)))) | ! [X44] : (((~ ! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p2(X44)) & (~ ! [X47] : (~ ! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | p2(X47) | ~ r1(X44, X47)) | ! [X50] : (! [X51] : (~ ! [X52] : (~ p2(X52) | ! [X53] : (p2(X53) | ~ r1(X52, X53)) | ~ r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) | ((~ ! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) | p2(X33)) & (~ ! [X56] : (~ ! [X57] : (~ p2(X57) | ! [X58] : (p2(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) | ! [X59] : (! [X60] : (~ ! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | p2(X60) | ~ r1(X59, X60)) | ~ r1(X33, X59)))) | ~ r1(X0, X33)) | ((~ ! [X63] : (~ p2(X63) | ! [X64] : (p2(X64) | ~ r1(X63, X64)) | ~ r1(X0, X63)) | p2(X0)) & (~ ! [X65] : (~ ! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) | p2(X65) | ~ r1(X0, X65)) | ! [X68] : (! [X69] : (~ ! [X70] : (~ p2(X70) | ! [X71] : (p2(X71) | ~ r1(X70, X71)) | ~ r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) | ~ ! [X72] : (~ ! [X73] : (~ p1(X73) | ! [X74] : (p1(X74) | ~ r1(X73, X74)) | ~ r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) | ! [X75] : (p1(X75) | ~ r1(X0, X75)) | ((~ ! [X76] : (~ ! [X77] : (~ p2(X77) | ! [X78] : (p2(X78) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) | ! [X79] : (p2(X79) | ~ r1(X0, X79))) & ~ ! [X80] : (~ ! [X81] : (~ p5(X81) | ~ r1(X80, X81)) | ~ r1(X0, X80))) | ~ ! [X82] : (~ ! [X83] : (~ p3(X83) | ! [X84] : (p3(X84) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) | ! [X85] : (p3(X85) | ~ r1(X0, X85)) | ((~ ! [X86] : (~ ! [X87] : (~ p2(X87) | ! [X88] : (p2(X88) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) | ! [X89] : (p2(X89) | ~ r1(X0, X89))) & ~ ! [X90] : (~ p5(X90) | ~ r1(X0, X90)))), inference(rectify, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ ! [X0] : (~ p5(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ p5(X1) | ~ r1(X0, X1)))), inference(negated_conjecture, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ ! [X0] : (~ p5(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ p5(X1) | ~ r1(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL660+1.005.p', main)).
fof(f30497, plain, (! [X0, X1] : (~ r1(X0, sK33(sK31(sK32))) | ~ sP1(X1) | ~ r1(X1, X0)) | ~ spl55_1692), inference(avatar_component_clause, [], [f30496])).
fof(f40369, plain, (~ spl55_37 | spl55_1785 | ~ spl55_35 | spl55_36 | ~ spl55_45 | ~ spl55_266), inference(avatar_split_clause, [], [f40368, f6996, f398, f361, f357, f32142, f366])).
fof(f366, plain, (spl55_37 <=> r1(sK32, sK44)), introduced(avatar_definition, [new_symbols(naming, [spl55_37])])).
fof(f32142, plain, (spl55_1785 <=> ! [X139] : (~ r1(sK52(sK44), X139) | p2(X139))), introduced(avatar_definition, [new_symbols(naming, [spl55_1785])])).
fof(f357, plain, (spl55_35 <=> ! [X29, X30] : (~ p2(X29) | ~ r1(sK44, X29) | ~ r1(X29, X30) | p2(X30))), introduced(avatar_definition, [new_symbols(naming, [spl55_35])])).
fof(f361, plain, (spl55_36 <=> p2(sK44)), introduced(avatar_definition, [new_symbols(naming, [spl55_36])])).
fof(f398, plain, (spl55_45 <=> ! [X41] : (r1(X41, sK52(X41)) | ~ r1(sK32, X41) | p2(X41))), introduced(avatar_definition, [new_symbols(naming, [spl55_45])])).
fof(f6996, plain, (spl55_266 <=> p2(sK52(sK44))), introduced(avatar_definition, [new_symbols(naming, [spl55_266])])).
fof(f40368, plain, (! [X17] : (~ r1(sK52(sK44), X17) | p2(X17) | ~ r1(sK32, sK44)) | (~ spl55_35 | spl55_36 | ~ spl55_45 | ~ spl55_266)), inference(subsumption_resolution, [], [f40367, f363])).
fof(f363, plain, (~ p2(sK44) | spl55_36), inference(avatar_component_clause, [], [f361])).
fof(f40367, plain, (! [X17] : (~ r1(sK52(sK44), X17) | p2(X17) | ~ r1(sK32, sK44) | p2(sK44)) | (~ spl55_35 | ~ spl55_45 | ~ spl55_266)), inference(subsumption_resolution, [], [f40209, f6997])).
fof(f6997, plain, (p2(sK52(sK44)) | ~ spl55_266), inference(avatar_component_clause, [], [f6996])).
fof(f40209, plain, (! [X17] : (~ p2(sK52(sK44)) | ~ r1(sK52(sK44), X17) | p2(X17) | ~ r1(sK32, sK44) | p2(sK44)) | (~ spl55_35 | ~ spl55_45)), inference(resolution, [], [f358, f399])).
fof(f399, plain, (! [X41] : (r1(X41, sK52(X41)) | ~ r1(sK32, X41) | p2(X41)) | ~ spl55_45), inference(avatar_component_clause, [], [f398])).
fof(f358, plain, (! [X30, X29] : (~ r1(sK44, X29) | ~ p2(X29) | ~ r1(X29, X30) | p2(X30)) | ~ spl55_35), inference(avatar_component_clause, [], [f357])).
fof(f40335, plain, (~ spl55_35 | ~ spl55_245 | ~ spl55_246 | ~ spl55_247 | spl55_1804), inference(avatar_contradiction_clause, [], [f40334])).
fof(f40334, plain, ($false | (~ spl55_35 | ~ spl55_245 | ~ spl55_246 | ~ spl55_247 | spl55_1804)), inference(subsumption_resolution, [], [f40313, f32507])).
fof(f32507, plain, (~ p2(sK30(sK44)) | spl55_1804), inference(avatar_component_clause, [], [f32506])).
fof(f32506, plain, (spl55_1804 <=> p2(sK30(sK44))), introduced(avatar_definition, [new_symbols(naming, [spl55_1804])])).
fof(f40313, plain, (p2(sK30(sK44)) | (~ spl55_35 | ~ spl55_245 | ~ spl55_246 | ~ spl55_247)), inference(resolution, [], [f40210, f6884])).
fof(f6884, plain, (r1(sK29(sK44), sK30(sK44)) | ~ spl55_246), inference(avatar_component_clause, [], [f6882])).
fof(f6882, plain, (spl55_246 <=> r1(sK29(sK44), sK30(sK44))), introduced(avatar_definition, [new_symbols(naming, [spl55_246])])).
fof(f40210, plain, (! [X0] : (~ r1(sK29(sK44), X0) | p2(X0)) | (~ spl55_35 | ~ spl55_245 | ~ spl55_247)), inference(subsumption_resolution, [], [f40192, f6878])).
fof(f6878, plain, (p2(sK29(sK44)) | ~ spl55_245), inference(avatar_component_clause, [], [f6876])).
fof(f6876, plain, (spl55_245 <=> p2(sK29(sK44))), introduced(avatar_definition, [new_symbols(naming, [spl55_245])])).
fof(f40192, plain, (! [X0] : (~ p2(sK29(sK44)) | ~ r1(sK29(sK44), X0) | p2(X0)) | (~ spl55_35 | ~ spl55_247)), inference(resolution, [], [f358, f6890])).
fof(f6890, plain, (r1(sK44, sK29(sK44)) | ~ spl55_247), inference(avatar_component_clause, [], [f6888])).
fof(f6888, plain, (spl55_247 <=> r1(sK44, sK29(sK44))), introduced(avatar_definition, [new_symbols(naming, [spl55_247])])).
fof(f40113, plain, (~ spl55_34 | spl55_257 | ~ spl55_2244), inference(avatar_contradiction_clause, [], [f40112])).
fof(f40112, plain, ($false | (~ spl55_34 | spl55_257 | ~ spl55_2244)), inference(subsumption_resolution, [], [f40111, f355])).
fof(f355, plain, (sP2(sK44) | ~ spl55_34), inference(avatar_component_clause, [], [f353])).
fof(f353, plain, (spl55_34 <=> sP2(sK44)), introduced(avatar_definition, [new_symbols(naming, [spl55_34])])).
fof(f40111, plain, (~ sP2(sK44) | (spl55_257 | ~ spl55_2244)), inference(duplicate_literal_removal, [], [f40109])).
fof(f40109, plain, (~ sP2(sK44) | ~ sP2(sK44) | (spl55_257 | ~ spl55_2244)), inference(resolution, [], [f38125, f130])).
fof(f130, plain, ! [X0] : (r1(X0, sK25(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f57])).
fof(f57, plain, ! [X0] : ((! [X1] : ((p2(sK23(X1)) & (~ p2(sK24(X1)) & r1(sK23(X1), sK24(X1))) & r1(X1, sK23(X1))) | p2(X1) | ~ r1(X0, X1)) & ((! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK26(X0), X6)) & ~ p2(sK26(X0)) & r1(sK25(X0), sK26(X0))) & r1(X0, sK25(X0)))) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK23, sK24, sK25, sK26])], [f52, f56, f55, f54, f53])).
fof(f53, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK23(X1)) & ? [X3] : (~ p2(X3) & r1(sK23(X1), X3)) & r1(X1, sK23(X1)))), introduced(choice_axiom, [])).
fof(f54, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK23(X1), X3)) => (~ p2(sK24(X1)) & r1(sK23(X1), sK24(X1)))), introduced(choice_axiom, [])).
fof(f55, plain, ! [X0] : (? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4)) => (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK25(X0), X5)) & r1(X0, sK25(X0)))), introduced(choice_axiom, [])).
fof(f56, plain, ! [X0] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK25(X0), X5)) => (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK26(X0), X6)) & ~ p2(sK26(X0)) & r1(sK25(X0), sK26(X0)))), introduced(choice_axiom, [])).
fof(f52, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4))) | ~ sP2(X0)), inference(rectify, [], [f51])).
fof(f51, plain, ! [X33] : ((! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59))) | ~ sP2(X33)), inference(nnf_transformation, [], [f11])).
fof(f38125, plain, (! [X0] : (~ r1(X0, sK25(sK44)) | ~ sP2(X0)) | (spl55_257 | ~ spl55_2244)), inference(subsumption_resolution, [], [f38124, f6957])).
fof(f6957, plain, (~ p2(sK25(sK44)) | spl55_257), inference(avatar_component_clause, [], [f6955])).
fof(f6955, plain, (spl55_257 <=> p2(sK25(sK44))), introduced(avatar_definition, [new_symbols(naming, [spl55_257])])).
fof(f38124, plain, (! [X0] : (p2(sK25(sK44)) | ~ r1(X0, sK25(sK44)) | ~ sP2(X0)) | ~ spl55_2244), inference(resolution, [], [f37872, f136])).
fof(f136, plain, ! [X0, X1] : (~ p2(sK24(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP2(X0)), inference(cnf_transformation, [], [f57])).
fof(f37872, plain, (p2(sK24(sK25(sK44))) | ~ spl55_2244), inference(avatar_component_clause, [], [f37870])).
fof(f37870, plain, (spl55_2244 <=> p2(sK24(sK25(sK44)))), introduced(avatar_definition, [new_symbols(naming, [spl55_2244])])).
fof(f37919, plain, (spl55_256 | ~ spl55_257 | ~ spl55_33 | ~ spl55_34 | spl55_255), inference(avatar_split_clause, [], [f37918, f6947, f353, f349, f6955, f6951])).
fof(f6951, plain, (spl55_256 <=> sP5(sK25(sK44))), introduced(avatar_definition, [new_symbols(naming, [spl55_256])])).
fof(f349, plain, (spl55_33 <=> ! [X26] : (~ p2(X26) | ~ r1(sK44, X26) | sP5(X26) | sP4(X26))), introduced(avatar_definition, [new_symbols(naming, [spl55_33])])).
fof(f6947, plain, (spl55_255 <=> sP4(sK25(sK44))), introduced(avatar_definition, [new_symbols(naming, [spl55_255])])).
fof(f37918, plain, (~ p2(sK25(sK44)) | sP5(sK25(sK44)) | (~ spl55_33 | ~ spl55_34 | spl55_255)), inference(subsumption_resolution, [], [f37917, f355])).
fof(f37917, plain, (~ p2(sK25(sK44)) | sP5(sK25(sK44)) | ~ sP2(sK44) | (~ spl55_33 | spl55_255)), inference(subsumption_resolution, [], [f6902, f6948])).
fof(f6948, plain, (~ sP4(sK25(sK44)) | spl55_255), inference(avatar_component_clause, [], [f6947])).
fof(f6902, plain, (~ p2(sK25(sK44)) | sP5(sK25(sK44)) | sP4(sK25(sK44)) | ~ sP2(sK44) | ~ spl55_33), inference(resolution, [], [f350, f130])).
fof(f350, plain, (! [X26] : (~ r1(sK44, X26) | ~ p2(X26) | sP5(X26) | sP4(X26)) | ~ spl55_33), inference(avatar_component_clause, [], [f349])).
fof(f37895, plain, (~ spl55_34 | ~ spl55_256), inference(avatar_contradiction_clause, [], [f37894])).
fof(f37894, plain, ($false | (~ spl55_34 | ~ spl55_256)), inference(subsumption_resolution, [], [f37893, f355])).
fof(f37893, plain, (~ sP2(sK44) | ~ spl55_256), inference(resolution, [], [f6953, f5413])).
fof(f5413, plain, ! [X1] : (~ sP5(sK25(X1)) | ~ sP2(X1)), inference(duplicate_literal_removal, [], [f5412])).
fof(f5412, plain, ! [X1] : (~ sP5(sK25(X1)) | ~ sP2(X1) | ~ sP5(sK25(X1)) | ~ sP2(X1)), inference(resolution, [], [f5320, f131])).
fof(f131, plain, ! [X0] : (r1(sK25(X0), sK26(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f57])).
fof(f5320, plain, ! [X0, X1] : (~ r1(X1, sK26(X0)) | ~ sP5(sK25(X0)) | ~ sP2(X0) | ~ sP5(X1)), inference(subsumption_resolution, [], [f5319, f132])).
fof(f132, plain, ! [X0] : (~ p2(sK26(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f57])).
fof(f5319, plain, ! [X0, X1] : (~ sP2(X0) | ~ sP5(sK25(X0)) | p2(sK26(X0)) | ~ r1(X1, sK26(X0)) | ~ sP5(X1)), inference(resolution, [], [f5280, f116])).
fof(f116, plain, ! [X0, X1] : (~ p2(sK15(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP5(X0)), inference(cnf_transformation, [], [f38])).
fof(f38, plain, ! [X0] : (! [X1] : ((((p2(sK14(X1)) & (~ p2(sK15(X1)) & r1(sK14(X1), sK15(X1))) & r1(X1, sK14(X1))) | p2(X1)) & ((! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(sK16(X1), X5)) & ~ p2(sK16(X1)) & r1(X1, sK16(X1))) | sP3(X1))) | ~ r1(X0, X1)) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14, sK15, sK16])], [f34, f37, f36, f35])).
fof(f35, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK14(X1)) & ? [X3] : (~ p2(X3) & r1(sK14(X1), X3)) & r1(X1, sK14(X1)))), introduced(choice_axiom, [])).
fof(f36, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK14(X1), X3)) => (~ p2(sK15(X1)) & r1(sK14(X1), sK15(X1)))), introduced(choice_axiom, [])).
fof(f37, plain, ! [X1] : (? [X4] : (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(X4, X5)) & ~ p2(X4) & r1(X1, X4)) => (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(sK16(X1), X5)) & ~ p2(sK16(X1)) & r1(X1, sK16(X1)))), introduced(choice_axiom, [])).
fof(f34, plain, ! [X0] : (! [X1] : (((? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1)) & (? [X4] : (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(X4, X5)) & ~ p2(X4) & r1(X1, X4)) | sP3(X1))) | ~ r1(X0, X1)) | ~ sP5(X0)), inference(rectify, [], [f33])).
fof(f33, plain, ! [X34] : (! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | sP3(X44))) | ~ r1(X34, X44)) | ~ sP5(X34)), inference(nnf_transformation, [], [f14])).
fof(f5280, plain, ! [X0] : (p2(sK15(sK26(X0))) | ~ sP2(X0) | ~ sP5(sK25(X0))), inference(duplicate_literal_removal, [], [f5260])).
fof(f5260, plain, ! [X0] : (~ sP2(X0) | p2(sK15(sK26(X0))) | ~ sP5(sK25(X0)) | ~ sP5(sK25(X0)) | ~ sP2(X0)), inference(resolution, [], [f2907, f1084])).
fof(f1084, plain, ! [X13] : (r1(sK14(sK26(X13)), sK15(sK26(X13))) | ~ sP5(sK25(X13)) | ~ sP2(X13)), inference(subsumption_resolution, [], [f1072, f132])).
fof(f1072, plain, ! [X13] : (p2(sK26(X13)) | r1(sK14(sK26(X13)), sK15(sK26(X13))) | ~ sP5(sK25(X13)) | ~ sP2(X13)), inference(resolution, [], [f115, f131])).
fof(f115, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK14(X1), sK15(X1)) | ~ sP5(X0)), inference(cnf_transformation, [], [f38])).
fof(f2907, plain, ! [X0, X1] : (~ r1(sK14(sK26(X0)), X1) | ~ sP2(X0) | p2(X1) | ~ sP5(sK25(X0))), inference(subsumption_resolution, [], [f2906, f790])).
fof(f790, plain, ! [X13] : (p2(sK14(sK26(X13))) | ~ sP5(sK25(X13)) | ~ sP2(X13)), inference(subsumption_resolution, [], [f779, f132])).
fof(f779, plain, ! [X13] : (p2(sK26(X13)) | p2(sK14(sK26(X13))) | ~ sP5(sK25(X13)) | ~ sP2(X13)), inference(resolution, [], [f117, f131])).
fof(f117, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK14(X1)) | ~ sP5(X0)), inference(cnf_transformation, [], [f38])).
fof(f2906, plain, ! [X0, X1] : (~ sP5(sK25(X0)) | ~ sP2(X0) | p2(X1) | ~ r1(sK14(sK26(X0)), X1) | ~ p2(sK14(sK26(X0)))), inference(duplicate_literal_removal, [], [f2875])).
fof(f2875, plain, ! [X0, X1] : (~ sP5(sK25(X0)) | ~ sP2(X0) | p2(X1) | ~ r1(sK14(sK26(X0)), X1) | ~ p2(sK14(sK26(X0))) | ~ sP2(X0)), inference(resolution, [], [f942, f133])).
fof(f133, plain, ! [X6, X0, X7] : (~ r1(sK26(X0), X6) | p2(X7) | ~ r1(X6, X7) | ~ p2(X6) | ~ sP2(X0)), inference(cnf_transformation, [], [f57])).
fof(f942, plain, ! [X13] : (r1(sK26(X13), sK14(sK26(X13))) | ~ sP5(sK25(X13)) | ~ sP2(X13)), inference(subsumption_resolution, [], [f931, f132])).
fof(f931, plain, ! [X13] : (p2(sK26(X13)) | r1(sK26(X13), sK14(sK26(X13))) | ~ sP5(sK25(X13)) | ~ sP2(X13)), inference(resolution, [], [f114, f131])).
fof(f114, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK14(X1)) | ~ sP5(X0)), inference(cnf_transformation, [], [f38])).
fof(f6953, plain, (sP5(sK25(sK44)) | ~ spl55_256), inference(avatar_component_clause, [], [f6951])).
fof(f37889, plain, (~ spl55_34 | ~ spl55_255), inference(avatar_contradiction_clause, [], [f37888])).
fof(f37888, plain, ($false | (~ spl55_34 | ~ spl55_255)), inference(subsumption_resolution, [], [f37887, f355])).
fof(f37887, plain, (~ sP2(sK44) | ~ spl55_255), inference(resolution, [], [f6949, f6058])).
fof(f6058, plain, ! [X1] : (~ sP4(sK25(X1)) | ~ sP2(X1)), inference(duplicate_literal_removal, [], [f6057])).
fof(f6057, plain, ! [X1] : (~ sP4(sK25(X1)) | ~ sP2(X1) | ~ sP4(sK25(X1)) | ~ sP2(X1)), inference(resolution, [], [f6055, f131])).
fof(f6055, plain, ! [X0, X1] : (~ r1(X1, sK26(X0)) | ~ sP4(sK25(X0)) | ~ sP2(X0) | ~ sP4(X1)), inference(subsumption_resolution, [], [f6054, f132])).
fof(f6054, plain, ! [X0, X1] : (~ sP2(X0) | ~ sP4(sK25(X0)) | p2(sK26(X0)) | ~ r1(X1, sK26(X0)) | ~ sP4(X1)), inference(resolution, [], [f6051, f124])).
fof(f124, plain, ! [X0, X1] : (~ p2(sK18(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP4(X0)), inference(cnf_transformation, [], [f45])).
fof(f45, plain, ! [X0] : ((! [X1] : ((p2(sK17(X1)) & (~ p2(sK18(X1)) & r1(sK17(X1), sK18(X1))) & r1(X1, sK17(X1))) | p2(X1) | ~ r1(X0, X1)) & ((! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK20(X0), X6)) & ~ p2(sK20(X0)) & r1(sK19(X0), sK20(X0))) & r1(X0, sK19(X0)))) | ~ sP4(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK17, sK18, sK19, sK20])], [f40, f44, f43, f42, f41])).
fof(f41, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK17(X1)) & ? [X3] : (~ p2(X3) & r1(sK17(X1), X3)) & r1(X1, sK17(X1)))), introduced(choice_axiom, [])).
fof(f42, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK17(X1), X3)) => (~ p2(sK18(X1)) & r1(sK17(X1), sK18(X1)))), introduced(choice_axiom, [])).
fof(f43, plain, ! [X0] : (? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4)) => (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK19(X0), X5)) & r1(X0, sK19(X0)))), introduced(choice_axiom, [])).
fof(f44, plain, ! [X0] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK19(X0), X5)) => (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK20(X0), X6)) & ~ p2(sK20(X0)) & r1(sK19(X0), sK20(X0)))), introduced(choice_axiom, [])).
fof(f40, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4))) | ~ sP4(X0)), inference(rectify, [], [f39])).
fof(f39, plain, ! [X34] : ((! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40))) | ~ sP4(X34)), inference(nnf_transformation, [], [f13])).
fof(f6051, plain, ! [X0] : (p2(sK18(sK26(X0))) | ~ sP2(X0) | ~ sP4(sK25(X0))), inference(duplicate_literal_removal, [], [f6031])).
fof(f6031, plain, ! [X0] : (~ sP2(X0) | p2(sK18(sK26(X0))) | ~ sP4(sK25(X0)) | ~ sP4(sK25(X0)) | ~ sP2(X0)), inference(resolution, [], [f3012, f1112])).
fof(f1112, plain, ! [X13] : (r1(sK17(sK26(X13)), sK18(sK26(X13))) | ~ sP4(sK25(X13)) | ~ sP2(X13)), inference(subsumption_resolution, [], [f1099, f132])).
fof(f1099, plain, ! [X13] : (p2(sK26(X13)) | r1(sK17(sK26(X13)), sK18(sK26(X13))) | ~ sP4(sK25(X13)) | ~ sP2(X13)), inference(resolution, [], [f123, f131])).
fof(f123, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK17(X1), sK18(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f45])).
fof(f3012, plain, ! [X0, X1] : (~ r1(sK17(sK26(X0)), X1) | ~ sP2(X0) | p2(X1) | ~ sP4(sK25(X0))), inference(subsumption_resolution, [], [f3011, f831])).
fof(f831, plain, ! [X13] : (p2(sK17(sK26(X13))) | ~ sP4(sK25(X13)) | ~ sP2(X13)), inference(subsumption_resolution, [], [f819, f132])).
fof(f819, plain, ! [X13] : (p2(sK26(X13)) | p2(sK17(sK26(X13))) | ~ sP4(sK25(X13)) | ~ sP2(X13)), inference(resolution, [], [f125, f131])).
fof(f125, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK17(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f45])).
fof(f3011, plain, ! [X0, X1] : (~ sP4(sK25(X0)) | ~ sP2(X0) | p2(X1) | ~ r1(sK17(sK26(X0)), X1) | ~ p2(sK17(sK26(X0)))), inference(duplicate_literal_removal, [], [f2980])).
fof(f2980, plain, ! [X0, X1] : (~ sP4(sK25(X0)) | ~ sP2(X0) | p2(X1) | ~ r1(sK17(sK26(X0)), X1) | ~ p2(sK17(sK26(X0))) | ~ sP2(X0)), inference(resolution, [], [f969, f133])).
fof(f969, plain, ! [X13] : (r1(sK26(X13), sK17(sK26(X13))) | ~ sP4(sK25(X13)) | ~ sP2(X13)), inference(subsumption_resolution, [], [f957, f132])).
fof(f957, plain, ! [X13] : (p2(sK26(X13)) | r1(sK26(X13), sK17(sK26(X13))) | ~ sP4(sK25(X13)) | ~ sP2(X13)), inference(resolution, [], [f122, f131])).
fof(f122, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK17(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f45])).
fof(f6949, plain, (sP4(sK25(sK44)) | ~ spl55_255), inference(avatar_component_clause, [], [f6947])).
fof(f37873, plain, (spl55_257 | spl55_2244 | ~ spl55_34 | ~ spl55_1897), inference(avatar_split_clause, [], [f37868, f33029, f353, f37870, f6955])).
fof(f33029, plain, (spl55_1897 <=> ! [X4] : (p2(X4) | ~ r1(sK23(sK25(sK44)), X4))), introduced(avatar_definition, [new_symbols(naming, [spl55_1897])])).
fof(f37868, plain, (p2(sK24(sK25(sK44))) | p2(sK25(sK44)) | (~ spl55_34 | ~ spl55_1897)), inference(subsumption_resolution, [], [f35039, f355])).
fof(f35039, plain, (p2(sK24(sK25(sK44))) | p2(sK25(sK44)) | ~ sP2(sK44) | ~ spl55_1897), inference(resolution, [], [f33030, f1159])).
fof(f1159, plain, ! [X4] : (r1(sK23(sK25(X4)), sK24(sK25(X4))) | p2(sK25(X4)) | ~ sP2(X4)), inference(duplicate_literal_removal, [], [f1142])).
fof(f1142, plain, ! [X4] : (p2(sK25(X4)) | r1(sK23(sK25(X4)), sK24(sK25(X4))) | ~ sP2(X4) | ~ sP2(X4)), inference(resolution, [], [f135, f130])).
fof(f135, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK23(X1), sK24(X1)) | ~ sP2(X0)), inference(cnf_transformation, [], [f57])).
fof(f33030, plain, (! [X4] : (~ r1(sK23(sK25(sK44)), X4) | p2(X4)) | ~ spl55_1897), inference(avatar_component_clause, [], [f33029])).
fof(f33328, plain, (spl55_36 | ~ spl55_37 | ~ spl55_39 | ~ spl55_1804), inference(avatar_contradiction_clause, [], [f33327])).
fof(f33327, plain, ($false | (spl55_36 | ~ spl55_37 | ~ spl55_39 | ~ spl55_1804)), inference(subsumption_resolution, [], [f33326, f376])).
fof(f376, plain, (sP0(sK32) | ~ spl55_39), inference(avatar_component_clause, [], [f374])).
fof(f374, plain, (spl55_39 <=> sP0(sK32)), introduced(avatar_definition, [new_symbols(naming, [spl55_39])])).
fof(f33326, plain, (~ sP0(sK32) | (spl55_36 | ~ spl55_37 | ~ spl55_1804)), inference(resolution, [], [f32906, f368])).
fof(f368, plain, (r1(sK32, sK44) | ~ spl55_37), inference(avatar_component_clause, [], [f366])).
fof(f32906, plain, (! [X0] : (~ r1(X0, sK44) | ~ sP0(X0)) | (spl55_36 | ~ spl55_1804)), inference(subsumption_resolution, [], [f32905, f363])).
fof(f32905, plain, (! [X0] : (p2(sK44) | ~ r1(X0, sK44) | ~ sP0(X0)) | ~ spl55_1804), inference(resolution, [], [f32508, f146])).
fof(f146, plain, ! [X0, X1] : (~ p2(sK30(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f68])).
fof(f68, plain, ! [X0] : ((! [X1] : ((p2(sK29(X1)) & (~ p2(sK30(X1)) & r1(sK29(X1), sK30(X1))) & r1(X1, sK29(X1))) | p2(X1) | ~ r1(X0, X1)) & (~ p2(sK31(X0)) & r1(X0, sK31(X0)))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK29, sK30, sK31])], [f64, f67, f66, f65])).
fof(f65, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK29(X1)) & ? [X3] : (~ p2(X3) & r1(sK29(X1), X3)) & r1(X1, sK29(X1)))), introduced(choice_axiom, [])).
fof(f66, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK29(X1), X3)) => (~ p2(sK30(X1)) & r1(sK29(X1), sK30(X1)))), introduced(choice_axiom, [])).
fof(f67, plain, ! [X0] : (? [X4] : (~ p2(X4) & r1(X0, X4)) => (~ p2(sK31(X0)) & r1(X0, sK31(X0)))), introduced(choice_axiom, [])).
fof(f64, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (~ p2(X4) & r1(X0, X4))) | ~ sP0(X0)), inference(rectify, [], [f63])).
fof(f63, plain, ! [X0] : ((! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79))) | ~ sP0(X0)), inference(nnf_transformation, [], [f9])).
fof(f32508, plain, (p2(sK30(sK44)) | ~ spl55_1804), inference(avatar_component_clause, [], [f32506])).
fof(f33161, plain, (~ spl55_34 | spl55_257 | spl55_1896), inference(avatar_contradiction_clause, [], [f33160])).
fof(f33160, plain, ($false | (~ spl55_34 | spl55_257 | spl55_1896)), inference(subsumption_resolution, [], [f33159, f355])).
fof(f33159, plain, (~ sP2(sK44) | (spl55_257 | spl55_1896)), inference(subsumption_resolution, [], [f33157, f6957])).
fof(f33157, plain, (p2(sK25(sK44)) | ~ sP2(sK44) | spl55_1896), inference(resolution, [], [f33027, f852])).
fof(f852, plain, ! [X4] : (p2(sK23(sK25(X4))) | p2(sK25(X4)) | ~ sP2(X4)), inference(duplicate_literal_removal, [], [f837])).
fof(f837, plain, ! [X4] : (p2(sK25(X4)) | p2(sK23(sK25(X4))) | ~ sP2(X4) | ~ sP2(X4)), inference(resolution, [], [f137, f130])).
fof(f137, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK23(X1)) | ~ sP2(X0)), inference(cnf_transformation, [], [f57])).
fof(f33027, plain, (~ p2(sK23(sK25(sK44))) | spl55_1896), inference(avatar_component_clause, [], [f33025])).
fof(f33025, plain, (spl55_1896 <=> p2(sK23(sK25(sK44)))), introduced(avatar_definition, [new_symbols(naming, [spl55_1896])])).
fof(f33031, plain, (~ spl55_1896 | spl55_1897 | ~ spl55_34 | spl55_257 | ~ spl55_286), inference(avatar_split_clause, [], [f33023, f7133, f6955, f353, f33029, f33025])).
fof(f7133, plain, (spl55_286 <=> ! [X18, X19] : (~ p2(X18) | p2(X19) | ~ r1(X18, X19) | ~ r1(sK25(sK44), X18))), introduced(avatar_definition, [new_symbols(naming, [spl55_286])])).
fof(f33023, plain, (! [X4] : (p2(X4) | ~ r1(sK23(sK25(sK44)), X4) | ~ p2(sK23(sK25(sK44)))) | (~ spl55_34 | spl55_257 | ~ spl55_286)), inference(subsumption_resolution, [], [f33022, f355])).
fof(f33022, plain, (! [X4] : (p2(X4) | ~ r1(sK23(sK25(sK44)), X4) | ~ p2(sK23(sK25(sK44))) | ~ sP2(sK44)) | (spl55_257 | ~ spl55_286)), inference(subsumption_resolution, [], [f32976, f6957])).
fof(f32976, plain, (! [X4] : (p2(X4) | ~ r1(sK23(sK25(sK44)), X4) | ~ p2(sK23(sK25(sK44))) | p2(sK25(sK44)) | ~ sP2(sK44)) | ~ spl55_286), inference(resolution, [], [f7134, f990])).
fof(f990, plain, ! [X4] : (r1(sK25(X4), sK23(sK25(X4))) | p2(sK25(X4)) | ~ sP2(X4)), inference(duplicate_literal_removal, [], [f975])).
fof(f975, plain, ! [X4] : (p2(sK25(X4)) | r1(sK25(X4), sK23(sK25(X4))) | ~ sP2(X4) | ~ sP2(X4)), inference(resolution, [], [f134, f130])).
fof(f134, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK23(X1)) | ~ sP2(X0)), inference(cnf_transformation, [], [f57])).
fof(f7134, plain, (! [X19, X18] : (~ r1(sK25(sK44), X18) | p2(X19) | ~ r1(X18, X19) | ~ p2(X18)) | ~ spl55_286), inference(avatar_component_clause, [], [f7133])).
fof(f32913, plain, (~ spl55_34 | spl55_286 | ~ spl55_32), inference(avatar_split_clause, [], [f32912, f345, f7133, f353])).
fof(f345, plain, (spl55_32 <=> ! [X27, X26, X28] : (~ p2(X27) | ~ r1(sK44, X26) | sP5(X26) | sP4(X26) | ~ r1(X26, X27) | ~ r1(X27, X28) | p2(X28))), introduced(avatar_definition, [new_symbols(naming, [spl55_32])])).
fof(f32912, plain, (! [X19, X18] : (~ p2(X18) | ~ r1(sK25(sK44), X18) | ~ r1(X18, X19) | p2(X19) | ~ sP2(sK44)) | ~ spl55_32), inference(subsumption_resolution, [], [f32911, f6058])).
fof(f32911, plain, (! [X19, X18] : (~ p2(X18) | sP4(sK25(sK44)) | ~ r1(sK25(sK44), X18) | ~ r1(X18, X19) | p2(X19) | ~ sP2(sK44)) | ~ spl55_32), inference(subsumption_resolution, [], [f32326, f5413])).
fof(f32326, plain, (! [X19, X18] : (~ p2(X18) | sP5(sK25(sK44)) | sP4(sK25(sK44)) | ~ r1(sK25(sK44), X18) | ~ r1(X18, X19) | p2(X19) | ~ sP2(sK44)) | ~ spl55_32), inference(resolution, [], [f346, f130])).
fof(f346, plain, (! [X28, X26, X27] : (~ r1(sK44, X26) | ~ p2(X27) | sP5(X26) | sP4(X26) | ~ r1(X26, X27) | ~ r1(X27, X28) | p2(X28)) | ~ spl55_32), inference(avatar_component_clause, [], [f345])).
fof(f32910, plain, (spl55_36 | ~ spl55_37 | ~ spl55_43 | ~ spl55_44 | ~ spl55_1785), inference(avatar_contradiction_clause, [], [f32909])).
fof(f32909, plain, ($false | (spl55_36 | ~ spl55_37 | ~ spl55_43 | ~ spl55_44 | ~ spl55_1785)), inference(subsumption_resolution, [], [f32908, f363])).
fof(f32908, plain, (p2(sK44) | (spl55_36 | ~ spl55_37 | ~ spl55_43 | ~ spl55_44 | ~ spl55_1785)), inference(subsumption_resolution, [], [f32907, f368])).
fof(f32907, plain, (~ r1(sK32, sK44) | p2(sK44) | (spl55_36 | ~ spl55_37 | ~ spl55_43 | ~ spl55_44 | ~ spl55_1785)), inference(resolution, [], [f32852, f391])).
fof(f391, plain, (! [X41] : (~ p2(sK53(X41)) | ~ r1(sK32, X41) | p2(X41)) | ~ spl55_43), inference(avatar_component_clause, [], [f390])).
fof(f390, plain, (spl55_43 <=> ! [X41] : (~ p2(sK53(X41)) | ~ r1(sK32, X41) | p2(X41))), introduced(avatar_definition, [new_symbols(naming, [spl55_43])])).
fof(f32852, plain, (p2(sK53(sK44)) | (spl55_36 | ~ spl55_37 | ~ spl55_44 | ~ spl55_1785)), inference(subsumption_resolution, [], [f32851, f363])).
fof(f32851, plain, (p2(sK53(sK44)) | p2(sK44) | (~ spl55_37 | ~ spl55_44 | ~ spl55_1785)), inference(subsumption_resolution, [], [f32833, f368])).
fof(f32833, plain, (p2(sK53(sK44)) | ~ r1(sK32, sK44) | p2(sK44) | (~ spl55_44 | ~ spl55_1785)), inference(resolution, [], [f32143, f395])).
fof(f395, plain, (! [X41] : (r1(sK52(X41), sK53(X41)) | ~ r1(sK32, X41) | p2(X41)) | ~ spl55_44), inference(avatar_component_clause, [], [f394])).
fof(f394, plain, (spl55_44 <=> ! [X41] : (r1(sK52(X41), sK53(X41)) | ~ r1(sK32, X41) | p2(X41))), introduced(avatar_definition, [new_symbols(naming, [spl55_44])])).
fof(f32143, plain, (! [X139] : (~ r1(sK52(sK44), X139) | p2(X139)) | ~ spl55_1785), inference(avatar_component_clause, [], [f32142])).
fof(f32316, plain, (spl55_266 | spl55_36 | ~ spl55_37 | ~ spl55_42), inference(avatar_split_clause, [], [f32315, f386, f366, f361, f6996])).
fof(f386, plain, (spl55_42 <=> ! [X41] : (p2(sK52(X41)) | ~ r1(sK32, X41) | p2(X41))), introduced(avatar_definition, [new_symbols(naming, [spl55_42])])).
fof(f32315, plain, (p2(sK52(sK44)) | (spl55_36 | ~ spl55_37 | ~ spl55_42)), inference(subsumption_resolution, [], [f32166, f363])).
fof(f32166, plain, (p2(sK52(sK44)) | p2(sK44) | (~ spl55_37 | ~ spl55_42)), inference(resolution, [], [f368, f387])).
fof(f387, plain, (! [X41] : (~ r1(sK32, X41) | p2(sK52(X41)) | p2(X41)) | ~ spl55_42), inference(avatar_component_clause, [], [f386])).
fof(f32151, plain, (spl55_85 | ~ spl55_31 | ~ spl55_39 | spl55_493), inference(avatar_split_clause, [], [f32150, f15091, f374, f341, f649])).
fof(f649, plain, (spl55_85 <=> sP1(sK32)), introduced(avatar_definition, [new_symbols(naming, [spl55_85])])).
fof(f341, plain, (spl55_31 <=> sP6(sK32)), introduced(avatar_definition, [new_symbols(naming, [spl55_31])])).
fof(f15091, plain, (spl55_493 <=> p2(sK30(sK13(sK32)))), introduced(avatar_definition, [new_symbols(naming, [spl55_493])])).
fof(f32150, plain, (~ sP6(sK32) | sP1(sK32) | (~ spl55_39 | spl55_493)), inference(subsumption_resolution, [], [f31109, f376])).
fof(f31109, plain, (~ sP6(sK32) | sP1(sK32) | ~ sP0(sK32) | spl55_493), inference(resolution, [], [f26780, f15092])).
fof(f15092, plain, (~ p2(sK30(sK13(sK32))) | spl55_493), inference(avatar_component_clause, [], [f15091])).
fof(f26780, plain, ! [X0] : (p2(sK30(sK13(X0))) | ~ sP6(X0) | sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f26760])).
fof(f26760, plain, ! [X0] : (sP1(X0) | ~ sP6(X0) | p2(sK30(sK13(X0))) | ~ sP0(X0) | ~ sP0(X0) | sP1(X0) | ~ sP6(X0)), inference(resolution, [], [f3866, f1189])).
fof(f1189, plain, ! [X2] : (r1(sK29(sK13(X2)), sK30(sK13(X2))) | ~ sP0(X2) | sP1(X2) | ~ sP6(X2)), inference(subsumption_resolution, [], [f1169, f105])).
fof(f105, plain, ! [X0] : (~ p2(sK13(X0)) | sP1(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : ((((p2(sK11(X0)) & (~ p2(sK12(X0)) & r1(sK11(X0), sK12(X0))) & r1(X0, sK11(X0))) | p2(X0)) & ((! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(sK13(X0), X4)) & ~ p2(sK13(X0)) & r1(X0, sK13(X0))) | sP1(X0))) | ~ sP6(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK11, sK12, sK13])], [f28, f31, f30, f29])).
fof(f29, plain, ! [X0] : (? [X1] : (p2(X1) & ? [X2] : (~ p2(X2) & r1(X1, X2)) & r1(X0, X1)) => (p2(sK11(X0)) & ? [X2] : (~ p2(X2) & r1(sK11(X0), X2)) & r1(X0, sK11(X0)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X0] : (? [X2] : (~ p2(X2) & r1(sK11(X0), X2)) => (~ p2(sK12(X0)) & r1(sK11(X0), sK12(X0)))), introduced(choice_axiom, [])).
fof(f31, plain, ! [X0] : (? [X3] : (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) & ~ p2(X3) & r1(X0, X3)) => (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(sK13(X0), X4)) & ~ p2(sK13(X0)) & r1(X0, sK13(X0)))), introduced(choice_axiom, [])).
fof(f28, plain, ! [X0] : (((? [X1] : (p2(X1) & ? [X2] : (~ p2(X2) & r1(X1, X2)) & r1(X0, X1)) | p2(X0)) & (? [X3] : (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) & ~ p2(X3) & r1(X0, X3)) | sP1(X0))) | ~ sP6(X0)), inference(rectify, [], [f27])).
fof(f27, plain, ! [X0] : (((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | sP1(X0))) | ~ sP6(X0)), inference(nnf_transformation, [], [f15])).
fof(f1169, plain, ! [X2] : (p2(sK13(X2)) | r1(sK29(sK13(X2)), sK30(sK13(X2))) | ~ sP0(X2) | sP1(X2) | ~ sP6(X2)), inference(resolution, [], [f145, f104])).
fof(f104, plain, ! [X0] : (r1(X0, sK13(X0)) | sP1(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f32])).
fof(f145, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK29(X1), sK30(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f68])).
fof(f3866, plain, ! [X0, X1] : (~ r1(sK29(sK13(X0)), X1) | sP1(X0) | ~ sP6(X0) | p2(X1) | ~ sP0(X0)), inference(subsumption_resolution, [], [f3865, f880])).
fof(f880, plain, ! [X2] : (p2(sK29(sK13(X2))) | ~ sP0(X2) | sP1(X2) | ~ sP6(X2)), inference(subsumption_resolution, [], [f862, f105])).
fof(f862, plain, ! [X2] : (p2(sK13(X2)) | p2(sK29(sK13(X2))) | ~ sP0(X2) | sP1(X2) | ~ sP6(X2)), inference(resolution, [], [f147, f104])).
fof(f147, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK29(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f68])).
fof(f3865, plain, ! [X0, X1] : (~ sP0(X0) | sP1(X0) | ~ sP6(X0) | p2(X1) | ~ r1(sK29(sK13(X0)), X1) | ~ p2(sK29(sK13(X0)))), inference(duplicate_literal_removal, [], [f3830])).
fof(f3830, plain, ! [X0, X1] : (~ sP0(X0) | sP1(X0) | ~ sP6(X0) | p2(X1) | ~ r1(sK29(sK13(X0)), X1) | ~ p2(sK29(sK13(X0))) | sP1(X0) | ~ sP6(X0)), inference(resolution, [], [f1018, f106])).
fof(f106, plain, ! [X4, X0, X5] : (~ r1(sK13(X0), X4) | p2(X5) | ~ r1(X4, X5) | ~ p2(X4) | sP1(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f32])).
fof(f1018, plain, ! [X2] : (r1(sK13(X2), sK29(sK13(X2))) | ~ sP0(X2) | sP1(X2) | ~ sP6(X2)), inference(subsumption_resolution, [], [f1000, f105])).
fof(f1000, plain, ! [X2] : (p2(sK13(X2)) | r1(sK13(X2), sK29(sK13(X2))) | ~ sP0(X2) | sP1(X2) | ~ sP6(X2)), inference(resolution, [], [f144, f104])).
fof(f144, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK29(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f68])).
fof(f32137, plain, (~ spl55_31 | spl55_85 | ~ spl55_42 | ~ spl55_43 | ~ spl55_44 | ~ spl55_45), inference(avatar_split_clause, [], [f20829, f398, f394, f390, f386, f649, f341])).
fof(f20829, plain, (sP1(sK32) | ~ sP6(sK32) | (~ spl55_42 | ~ spl55_43 | ~ spl55_44 | ~ spl55_45)), inference(duplicate_literal_removal, [], [f20828])).
fof(f20828, plain, (sP1(sK32) | ~ sP6(sK32) | sP1(sK32) | ~ sP6(sK32) | (~ spl55_42 | ~ spl55_43 | ~ spl55_44 | ~ spl55_45)), inference(resolution, [], [f19916, f104])).
fof(f19916, plain, (! [X1] : (~ r1(sK32, sK13(X1)) | sP1(X1) | ~ sP6(X1)) | (~ spl55_42 | ~ spl55_43 | ~ spl55_44 | ~ spl55_45)), inference(subsumption_resolution, [], [f19915, f105])).
fof(f19915, plain, (! [X1] : (~ r1(sK32, sK13(X1)) | sP1(X1) | ~ sP6(X1) | p2(sK13(X1))) | (~ spl55_42 | ~ spl55_43 | ~ spl55_44 | ~ spl55_45)), inference(subsumption_resolution, [], [f19914, f387])).
fof(f19914, plain, (! [X1] : (~ r1(sK32, sK13(X1)) | ~ p2(sK52(sK13(X1))) | sP1(X1) | ~ sP6(X1) | p2(sK13(X1))) | (~ spl55_43 | ~ spl55_44 | ~ spl55_45)), inference(subsumption_resolution, [], [f19910, f391])).
fof(f19910, plain, (! [X1] : (p2(sK53(sK13(X1))) | ~ r1(sK32, sK13(X1)) | ~ p2(sK52(sK13(X1))) | sP1(X1) | ~ sP6(X1) | p2(sK13(X1))) | (~ spl55_44 | ~ spl55_45)), inference(duplicate_literal_removal, [], [f19891])).
fof(f19891, plain, (! [X1] : (p2(sK53(sK13(X1))) | ~ r1(sK32, sK13(X1)) | ~ p2(sK52(sK13(X1))) | sP1(X1) | ~ sP6(X1) | ~ r1(sK32, sK13(X1)) | p2(sK13(X1))) | (~ spl55_44 | ~ spl55_45)), inference(resolution, [], [f17129, f395])).
fof(f17129, plain, (! [X99, X98] : (~ r1(sK52(sK13(X98)), X99) | p2(X99) | ~ r1(sK32, sK13(X98)) | ~ p2(sK52(sK13(X98))) | sP1(X98) | ~ sP6(X98)) | ~ spl55_45), inference(subsumption_resolution, [], [f17072, f105])).
fof(f17072, plain, (! [X99, X98] : (~ r1(sK32, sK13(X98)) | p2(sK13(X98)) | p2(X99) | ~ r1(sK52(sK13(X98)), X99) | ~ p2(sK52(sK13(X98))) | sP1(X98) | ~ sP6(X98)) | ~ spl55_45), inference(resolution, [], [f399, f106])).
fof(f30506, plain, (spl55_80 | ~ spl55_399 | ~ spl55_1693), inference(avatar_contradiction_clause, [], [f30505])).
fof(f30505, plain, ($false | (spl55_80 | ~ spl55_399 | ~ spl55_1693)), inference(subsumption_resolution, [], [f30504, f13440])).
fof(f30504, plain, (~ r1(sK32, sK31(sK32)) | (spl55_80 | ~ spl55_1693)), inference(subsumption_resolution, [], [f30503, f623])).
fof(f30503, plain, (p2(sK31(sK32)) | ~ r1(sK32, sK31(sK32)) | ~ spl55_1693), inference(resolution, [], [f30501, f199])).
fof(f199, plain, ! [X1] : (~ p2(sK33(X1)) | p2(X1) | ~ r1(sK32, X1)), inference(cnf_transformation, [], [f93])).
fof(f30501, plain, (p2(sK33(sK31(sK32))) | ~ spl55_1693), inference(avatar_component_clause, [], [f30499])).
fof(f30499, plain, (spl55_1693 <=> p2(sK33(sK31(sK32)))), introduced(avatar_definition, [new_symbols(naming, [spl55_1693])])).
fof(f30502, plain, (spl55_1692 | spl55_1693 | ~ spl55_571 | ~ spl55_869), inference(avatar_split_clause, [], [f30494, f21090, f15960, f30499, f30496])).
fof(f15960, plain, (spl55_571 <=> r1(sK27(sK33(sK31(sK32))), sK28(sK33(sK31(sK32))))), introduced(avatar_definition, [new_symbols(naming, [spl55_571])])).
fof(f21090, plain, (spl55_869 <=> ! [X0] : (p2(X0) | ~ r1(sK27(sK33(sK31(sK32))), X0))), introduced(avatar_definition, [new_symbols(naming, [spl55_869])])).
fof(f30494, plain, (! [X0, X1] : (p2(sK33(sK31(sK32))) | ~ r1(X0, sK33(sK31(sK32))) | ~ r1(X1, X0) | ~ sP1(X1)) | (~ spl55_571 | ~ spl55_869)), inference(resolution, [], [f30230, f140])).
fof(f140, plain, ! [X2, X0, X1] : (~ p2(sK28(X2)) | p2(X2) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f62])).
fof(f62, plain, ! [X0] : (! [X1] : (! [X2] : ((p2(sK27(X2)) & (~ p2(sK28(X2)) & r1(sK27(X2), sK28(X2))) & r1(X2, sK27(X2))) | p2(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK27, sK28])], [f59, f61, f60])).
fof(f60, plain, ! [X2] : (? [X3] : (p2(X3) & ? [X4] : (~ p2(X4) & r1(X3, X4)) & r1(X2, X3)) => (p2(sK27(X2)) & ? [X4] : (~ p2(X4) & r1(sK27(X2), X4)) & r1(X2, sK27(X2)))), introduced(choice_axiom, [])).
fof(f61, plain, ! [X2] : (? [X4] : (~ p2(X4) & r1(sK27(X2), X4)) => (~ p2(sK28(X2)) & r1(sK27(X2), sK28(X2)))), introduced(choice_axiom, [])).
fof(f59, plain, ! [X0] : (! [X1] : (! [X2] : (? [X3] : (p2(X3) & ? [X4] : (~ p2(X4) & r1(X3, X4)) & r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(rectify, [], [f58])).
fof(f58, plain, ! [X0] : (! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)) | ~ sP1(X0)), inference(nnf_transformation, [], [f10])).
fof(f30230, plain, (p2(sK28(sK33(sK31(sK32)))) | (~ spl55_571 | ~ spl55_869)), inference(resolution, [], [f15962, f21091])).
fof(f21091, plain, (! [X0] : (~ r1(sK27(sK33(sK31(sK32))), X0) | p2(X0)) | ~ spl55_869), inference(avatar_component_clause, [], [f21090])).
fof(f15962, plain, (r1(sK27(sK33(sK31(sK32))), sK28(sK33(sK31(sK32)))) | ~ spl55_571), inference(avatar_component_clause, [], [f15960])).
fof(f29083, plain, (spl55_1062 | spl55_46 | ~ spl55_47 | ~ spl55_588 | ~ spl55_592 | ~ spl55_594), inference(avatar_split_clause, [], [f24876, f16949, f16937, f16915, f407, f402, f22956])).
fof(f22956, plain, (spl55_1062 <=> p2(sK28(sK33(sK54)))), introduced(avatar_definition, [new_symbols(naming, [spl55_1062])])).
fof(f402, plain, (spl55_46 <=> p2(sK54)), introduced(avatar_definition, [new_symbols(naming, [spl55_46])])).
fof(f407, plain, (spl55_47 <=> r1(sK32, sK54)), introduced(avatar_definition, [new_symbols(naming, [spl55_47])])).
fof(f16915, plain, (spl55_588 <=> p2(sK27(sK33(sK54)))), introduced(avatar_definition, [new_symbols(naming, [spl55_588])])).
fof(f16937, plain, (spl55_592 <=> r1(sK33(sK54), sK27(sK33(sK54)))), introduced(avatar_definition, [new_symbols(naming, [spl55_592])])).
fof(f16949, plain, (spl55_594 <=> r1(sK27(sK33(sK54)), sK28(sK33(sK54)))), introduced(avatar_definition, [new_symbols(naming, [spl55_594])])).
fof(f24876, plain, (p2(sK28(sK33(sK54))) | (spl55_46 | ~ spl55_47 | ~ spl55_588 | ~ spl55_592 | ~ spl55_594)), inference(resolution, [], [f22634, f16951])).
fof(f16951, plain, (r1(sK27(sK33(sK54)), sK28(sK33(sK54))) | ~ spl55_594), inference(avatar_component_clause, [], [f16949])).
fof(f22634, plain, (! [X0] : (~ r1(sK27(sK33(sK54)), X0) | p2(X0)) | (spl55_46 | ~ spl55_47 | ~ spl55_588 | ~ spl55_592)), inference(subsumption_resolution, [], [f22633, f409])).
fof(f409, plain, (r1(sK32, sK54) | ~ spl55_47), inference(avatar_component_clause, [], [f407])).
fof(f22633, plain, (! [X0] : (p2(X0) | ~ r1(sK27(sK33(sK54)), X0) | ~ r1(sK32, sK54)) | (spl55_46 | ~ spl55_588 | ~ spl55_592)), inference(subsumption_resolution, [], [f22632, f404])).
fof(f404, plain, (~ p2(sK54) | spl55_46), inference(avatar_component_clause, [], [f402])).
fof(f22632, plain, (! [X0] : (p2(X0) | ~ r1(sK27(sK33(sK54)), X0) | p2(sK54) | ~ r1(sK32, sK54)) | (~ spl55_588 | ~ spl55_592)), inference(subsumption_resolution, [], [f22564, f16917])).
fof(f16917, plain, (p2(sK27(sK33(sK54))) | ~ spl55_588), inference(avatar_component_clause, [], [f16915])).
fof(f22564, plain, (! [X0] : (p2(X0) | ~ r1(sK27(sK33(sK54)), X0) | ~ p2(sK27(sK33(sK54))) | p2(sK54) | ~ r1(sK32, sK54)) | ~ spl55_592), inference(resolution, [], [f16939, f200])).
fof(f200, plain, ! [X4, X3, X1] : (~ r1(sK33(X1), X3) | p2(X4) | ~ r1(X3, X4) | ~ p2(X3) | p2(X1) | ~ r1(sK32, X1)), inference(cnf_transformation, [], [f93])).
fof(f16939, plain, (r1(sK33(sK54), sK27(sK33(sK54))) | ~ spl55_592), inference(avatar_component_clause, [], [f16937])).
fof(f29082, plain, (spl55_46 | ~ spl55_47 | ~ spl55_1320), inference(avatar_contradiction_clause, [], [f29081])).
fof(f29081, plain, ($false | (spl55_46 | ~ spl55_47 | ~ spl55_1320)), inference(subsumption_resolution, [], [f29080, f409])).
fof(f29080, plain, (~ r1(sK32, sK54) | (spl55_46 | ~ spl55_1320)), inference(subsumption_resolution, [], [f29079, f404])).
fof(f29079, plain, (p2(sK54) | ~ r1(sK32, sK54) | ~ spl55_1320), inference(resolution, [], [f24416, f199])).
fof(f24416, plain, (p2(sK33(sK54)) | ~ spl55_1320), inference(avatar_component_clause, [], [f24414])).
fof(f24414, plain, (spl55_1320 <=> p2(sK33(sK54))), introduced(avatar_definition, [new_symbols(naming, [spl55_1320])])).
fof(f29078, plain, (spl55_586 | spl55_46 | ~ spl55_47 | ~ spl55_1319), inference(avatar_split_clause, [], [f29077, f24411, f407, f402, f16906])).
fof(f16906, plain, (spl55_586 <=> ! [X5] : (~ r1(X5, sK54) | ~ sP1(X5))), introduced(avatar_definition, [new_symbols(naming, [spl55_586])])).
fof(f24411, plain, (spl55_1319 <=> ! [X1, X0] : (~ r1(X0, sK33(sK54)) | ~ sP1(X1) | ~ r1(X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl55_1319])])).
fof(f29077, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK54)) | (spl55_46 | ~ spl55_47 | ~ spl55_1319)), inference(subsumption_resolution, [], [f29076, f409])).
fof(f29076, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK54) | ~ r1(sK32, sK54)) | (spl55_46 | ~ spl55_1319)), inference(subsumption_resolution, [], [f29074, f404])).
fof(f29074, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK54) | p2(sK54) | ~ r1(sK32, sK54)) | ~ spl55_1319), inference(resolution, [], [f24412, f198])).
fof(f24412, plain, (! [X0, X1] : (~ r1(X0, sK33(sK54)) | ~ sP1(X1) | ~ r1(X1, X0)) | ~ spl55_1319), inference(avatar_component_clause, [], [f24411])).
fof(f24417, plain, (spl55_1319 | spl55_1320 | ~ spl55_1062), inference(avatar_split_clause, [], [f24409, f22956, f24414, f24411])).
fof(f24409, plain, (! [X0, X1] : (p2(sK33(sK54)) | ~ r1(X0, sK33(sK54)) | ~ r1(X1, X0) | ~ sP1(X1)) | ~ spl55_1062), inference(resolution, [], [f22958, f140])).
fof(f22958, plain, (p2(sK28(sK33(sK54))) | ~ spl55_1062), inference(avatar_component_clause, [], [f22956])).
fof(f21092, plain, (~ spl55_399 | spl55_869 | spl55_80 | ~ spl55_289 | ~ spl55_380), inference(avatar_split_clause, [], [f17879, f11307, f7172, f622, f21090, f13439])).
fof(f7172, plain, (spl55_289 <=> p2(sK27(sK33(sK31(sK32))))), introduced(avatar_definition, [new_symbols(naming, [spl55_289])])).
fof(f11307, plain, (spl55_380 <=> r1(sK33(sK31(sK32)), sK27(sK33(sK31(sK32))))), introduced(avatar_definition, [new_symbols(naming, [spl55_380])])).
fof(f17879, plain, (! [X0] : (p2(X0) | ~ r1(sK27(sK33(sK31(sK32))), X0) | ~ r1(sK32, sK31(sK32))) | (spl55_80 | ~ spl55_289 | ~ spl55_380)), inference(subsumption_resolution, [], [f17878, f623])).
fof(f17878, plain, (! [X0] : (p2(X0) | ~ r1(sK27(sK33(sK31(sK32))), X0) | p2(sK31(sK32)) | ~ r1(sK32, sK31(sK32))) | (~ spl55_289 | ~ spl55_380)), inference(subsumption_resolution, [], [f17816, f7174])).
fof(f7174, plain, (p2(sK27(sK33(sK31(sK32)))) | ~ spl55_289), inference(avatar_component_clause, [], [f7172])).
fof(f17816, plain, (! [X0] : (p2(X0) | ~ r1(sK27(sK33(sK31(sK32))), X0) | ~ p2(sK27(sK33(sK31(sK32)))) | p2(sK31(sK32)) | ~ r1(sK32, sK31(sK32))) | ~ spl55_380), inference(resolution, [], [f11309, f200])).
fof(f11309, plain, (r1(sK33(sK31(sK32)), sK27(sK33(sK31(sK32)))) | ~ spl55_380), inference(avatar_component_clause, [], [f11307])).
fof(f20840, plain, (~ spl55_85 | ~ spl55_47 | ~ spl55_586), inference(avatar_split_clause, [], [f17313, f16906, f407, f649])).
fof(f17313, plain, (~ sP1(sK32) | (~ spl55_47 | ~ spl55_586)), inference(resolution, [], [f16907, f409])).
fof(f16907, plain, (! [X5] : (~ r1(X5, sK54) | ~ sP1(X5)) | ~ spl55_586), inference(avatar_component_clause, [], [f16906])).
fof(f16952, plain, (spl55_594 | spl55_586 | spl55_46 | ~ spl55_47), inference(avatar_split_clause, [], [f16947, f407, f402, f16906, f16949])).
fof(f16947, plain, (! [X13] : (~ r1(X13, sK54) | ~ sP1(X13) | r1(sK27(sK33(sK54)), sK28(sK33(sK54)))) | (spl55_46 | ~ spl55_47)), inference(subsumption_resolution, [], [f16825, f404])).
fof(f16825, plain, (! [X13] : (~ r1(X13, sK54) | ~ sP1(X13) | p2(sK54) | r1(sK27(sK33(sK54)), sK28(sK33(sK54)))) | ~ spl55_47), inference(resolution, [], [f409, f1571])).
fof(f1571, plain, ! [X12, X13] : (~ r1(sK32, X12) | ~ r1(X13, X12) | ~ sP1(X13) | p2(X12) | r1(sK27(sK33(X12)), sK28(sK33(X12)))), inference(subsumption_resolution, [], [f1553, f199])).
fof(f1553, plain, ! [X12, X13] : (p2(sK33(X12)) | r1(sK27(sK33(X12)), sK28(sK33(X12))) | ~ r1(X13, X12) | ~ sP1(X13) | p2(X12) | ~ r1(sK32, X12)), inference(resolution, [], [f139, f198])).
fof(f139, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | r1(sK27(X2), sK28(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f62])).
fof(f16940, plain, (spl55_592 | spl55_586 | spl55_46 | ~ spl55_47), inference(avatar_split_clause, [], [f16935, f407, f402, f16906, f16937])).
fof(f16935, plain, (! [X11] : (~ r1(X11, sK54) | ~ sP1(X11) | r1(sK33(sK54), sK27(sK33(sK54)))) | (spl55_46 | ~ spl55_47)), inference(subsumption_resolution, [], [f16823, f404])).
fof(f16823, plain, (! [X11] : (~ r1(X11, sK54) | ~ sP1(X11) | p2(sK54) | r1(sK33(sK54), sK27(sK33(sK54)))) | ~ spl55_47), inference(resolution, [], [f409, f1486])).
fof(f1486, plain, ! [X12, X13] : (~ r1(sK32, X12) | ~ r1(X13, X12) | ~ sP1(X13) | p2(X12) | r1(sK33(X12), sK27(sK33(X12)))), inference(subsumption_resolution, [], [f1468, f199])).
fof(f1468, plain, ! [X12, X13] : (p2(sK33(X12)) | r1(sK33(X12), sK27(sK33(X12))) | ~ r1(X13, X12) | ~ sP1(X13) | p2(X12) | ~ r1(sK32, X12)), inference(resolution, [], [f138, f198])).
fof(f138, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | r1(X2, sK27(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f62])).
fof(f16918, plain, (spl55_588 | spl55_586 | spl55_46 | ~ spl55_47), inference(avatar_split_clause, [], [f16913, f407, f402, f16906, f16915])).
fof(f16913, plain, (! [X7] : (~ r1(X7, sK54) | ~ sP1(X7) | p2(sK27(sK33(sK54)))) | (spl55_46 | ~ spl55_47)), inference(subsumption_resolution, [], [f16819, f404])).
fof(f16819, plain, (! [X7] : (~ r1(X7, sK54) | ~ sP1(X7) | p2(sK54) | p2(sK27(sK33(sK54)))) | ~ spl55_47), inference(resolution, [], [f409, f1407])).
fof(f1407, plain, ! [X12, X13] : (~ r1(sK32, X12) | ~ r1(X13, X12) | ~ sP1(X13) | p2(X12) | p2(sK27(sK33(X12)))), inference(subsumption_resolution, [], [f1389, f199])).
fof(f1389, plain, ! [X12, X13] : (p2(sK33(X12)) | p2(sK27(sK33(X12))) | ~ r1(X13, X12) | ~ sP1(X13) | p2(X12) | ~ r1(sK32, X12)), inference(resolution, [], [f141, f198])).
fof(f141, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | p2(sK27(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f62])).
fof(f16341, plain, (~ spl55_31 | ~ spl55_39 | spl55_85 | spl55_228 | ~ spl55_493), inference(avatar_contradiction_clause, [], [f16340])).
fof(f16340, plain, ($false | (~ spl55_31 | ~ spl55_39 | spl55_85 | spl55_228 | ~ spl55_493)), inference(subsumption_resolution, [], [f16339, f343])).
fof(f343, plain, (sP6(sK32) | ~ spl55_31), inference(avatar_component_clause, [], [f341])).
fof(f16339, plain, (~ sP6(sK32) | (~ spl55_39 | spl55_85 | spl55_228 | ~ spl55_493)), inference(subsumption_resolution, [], [f16338, f650])).
fof(f650, plain, (~ sP1(sK32) | spl55_85), inference(avatar_component_clause, [], [f649])).
fof(f16338, plain, (sP1(sK32) | ~ sP6(sK32) | (~ spl55_39 | spl55_228 | ~ spl55_493)), inference(subsumption_resolution, [], [f16335, f376])).
fof(f16335, plain, (~ sP0(sK32) | sP1(sK32) | ~ sP6(sK32) | (spl55_228 | ~ spl55_493)), inference(resolution, [], [f15540, f104])).
fof(f15540, plain, (! [X0] : (~ r1(X0, sK13(sK32)) | ~ sP0(X0)) | (spl55_228 | ~ spl55_493)), inference(subsumption_resolution, [], [f15539, f6750])).
fof(f6750, plain, (~ p2(sK13(sK32)) | spl55_228), inference(avatar_component_clause, [], [f6749])).
fof(f6749, plain, (spl55_228 <=> p2(sK13(sK32))), introduced(avatar_definition, [new_symbols(naming, [spl55_228])])).
fof(f15539, plain, (! [X0] : (p2(sK13(sK32)) | ~ r1(X0, sK13(sK32)) | ~ sP0(X0)) | ~ spl55_493), inference(resolution, [], [f15093, f146])).
fof(f15093, plain, (p2(sK30(sK13(sK32))) | ~ spl55_493), inference(avatar_component_clause, [], [f15091])).
fof(f15963, plain, (spl55_571 | spl55_290 | spl55_80 | ~ spl55_399), inference(avatar_split_clause, [], [f15958, f13439, f622, f7176, f15960])).
fof(f15958, plain, (! [X1] : (~ r1(X1, sK31(sK32)) | ~ sP1(X1) | r1(sK27(sK33(sK31(sK32))), sK28(sK33(sK31(sK32))))) | (spl55_80 | ~ spl55_399)), inference(subsumption_resolution, [], [f15932, f623])).
fof(f15932, plain, (! [X1] : (~ r1(X1, sK31(sK32)) | ~ sP1(X1) | p2(sK31(sK32)) | r1(sK27(sK33(sK31(sK32))), sK28(sK33(sK31(sK32))))) | ~ spl55_399), inference(resolution, [], [f1571, f13440])).
fof(f13948, plain, (~ spl55_31 | spl55_85 | ~ spl55_228), inference(avatar_contradiction_clause, [], [f13947])).
fof(f13947, plain, ($false | (~ spl55_31 | spl55_85 | ~ spl55_228)), inference(subsumption_resolution, [], [f13946, f343])).
fof(f13946, plain, (~ sP6(sK32) | (spl55_85 | ~ spl55_228)), inference(subsumption_resolution, [], [f13945, f650])).
fof(f13945, plain, (sP1(sK32) | ~ sP6(sK32) | ~ spl55_228), inference(resolution, [], [f6751, f105])).
fof(f6751, plain, (p2(sK13(sK32)) | ~ spl55_228), inference(avatar_component_clause, [], [f6749])).
fof(f13683, plain, (~ spl55_39 | spl55_399), inference(avatar_contradiction_clause, [], [f13682])).
fof(f13682, plain, ($false | (~ spl55_39 | spl55_399)), inference(subsumption_resolution, [], [f13681, f376])).
fof(f13681, plain, (~ sP0(sK32) | spl55_399), inference(resolution, [], [f13441, f142])).
fof(f142, plain, ! [X0] : (r1(X0, sK31(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f68])).
fof(f13441, plain, (~ r1(sK32, sK31(sK32)) | spl55_399), inference(avatar_component_clause, [], [f13439])).
fof(f11310, plain, (spl55_380 | spl55_290 | ~ spl55_39 | spl55_80), inference(avatar_split_clause, [], [f11305, f622, f374, f7176, f11307])).
fof(f11305, plain, (! [X14] : (~ r1(X14, sK31(sK32)) | ~ sP1(X14) | r1(sK33(sK31(sK32)), sK27(sK33(sK31(sK32))))) | (~ spl55_39 | spl55_80)), inference(subsumption_resolution, [], [f11304, f376])).
fof(f11304, plain, (! [X14] : (~ r1(X14, sK31(sK32)) | ~ sP1(X14) | r1(sK33(sK31(sK32)), sK27(sK33(sK31(sK32)))) | ~ sP0(sK32)) | spl55_80), inference(subsumption_resolution, [], [f11157, f623])).
fof(f11157, plain, ! [X14] : (~ r1(X14, sK31(sK32)) | ~ sP1(X14) | p2(sK31(sK32)) | r1(sK33(sK31(sK32)), sK27(sK33(sK31(sK32)))) | ~ sP0(sK32)), inference(resolution, [], [f1486, f142])).
fof(f11272, plain, (~ spl55_85 | ~ spl55_39 | ~ spl55_290), inference(avatar_split_clause, [], [f11271, f7176, f374, f649])).
fof(f11271, plain, (~ sP1(sK32) | (~ spl55_39 | ~ spl55_290)), inference(subsumption_resolution, [], [f7604, f376])).
fof(f7604, plain, (~ sP1(sK32) | ~ sP0(sK32) | ~ spl55_290), inference(resolution, [], [f7177, f142])).
fof(f7177, plain, (! [X14] : (~ r1(X14, sK31(sK32)) | ~ sP1(X14)) | ~ spl55_290), inference(avatar_component_clause, [], [f7176])).
fof(f7178, plain, (spl55_289 | spl55_290 | ~ spl55_39 | spl55_80), inference(avatar_split_clause, [], [f7170, f622, f374, f7176, f7172])).
fof(f7170, plain, (! [X14] : (~ r1(X14, sK31(sK32)) | ~ sP1(X14) | p2(sK27(sK33(sK31(sK32))))) | (~ spl55_39 | spl55_80)), inference(subsumption_resolution, [], [f7169, f376])).
fof(f7169, plain, (! [X14] : (~ r1(X14, sK31(sK32)) | ~ sP1(X14) | p2(sK27(sK33(sK31(sK32)))) | ~ sP0(sK32)) | spl55_80), inference(subsumption_resolution, [], [f7153, f623])).
fof(f7153, plain, ! [X14] : (~ r1(X14, sK31(sK32)) | ~ sP1(X14) | p2(sK31(sK32)) | p2(sK27(sK33(sK31(sK32)))) | ~ sP0(sK32)), inference(resolution, [], [f1407, f142])).
fof(f7138, plain, (~ spl55_39 | ~ spl55_80), inference(avatar_contradiction_clause, [], [f7137])).
fof(f7137, plain, ($false | (~ spl55_39 | ~ spl55_80)), inference(subsumption_resolution, [], [f7136, f376])).
fof(f7136, plain, (~ sP0(sK32) | ~ spl55_80), inference(resolution, [], [f624, f143])).
fof(f143, plain, ! [X0] : (~ p2(sK31(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f68])).
fof(f624, plain, (p2(sK31(sK32)) | ~ spl55_80), inference(avatar_component_clause, [], [f622])).
fof(f6891, plain, (spl55_247 | spl55_36 | ~ spl55_37 | ~ spl55_39), inference(avatar_split_clause, [], [f6886, f374, f366, f361, f6888])).
fof(f6886, plain, (p2(sK44) | r1(sK44, sK29(sK44)) | (~ spl55_37 | ~ spl55_39)), inference(subsumption_resolution, [], [f6826, f376])).
fof(f6826, plain, (p2(sK44) | r1(sK44, sK29(sK44)) | ~ sP0(sK32) | ~ spl55_37), inference(resolution, [], [f368, f144])).
fof(f6885, plain, (spl55_246 | spl55_36 | ~ spl55_37 | ~ spl55_39), inference(avatar_split_clause, [], [f6880, f374, f366, f361, f6882])).
fof(f6880, plain, (p2(sK44) | r1(sK29(sK44), sK30(sK44)) | (~ spl55_37 | ~ spl55_39)), inference(subsumption_resolution, [], [f6827, f376])).
fof(f6827, plain, (p2(sK44) | r1(sK29(sK44), sK30(sK44)) | ~ sP0(sK32) | ~ spl55_37), inference(resolution, [], [f368, f145])).
fof(f6879, plain, (spl55_245 | spl55_36 | ~ spl55_37 | ~ spl55_39), inference(avatar_split_clause, [], [f6874, f374, f366, f361, f6876])).
fof(f6874, plain, (p2(sK44) | p2(sK29(sK44)) | (~ spl55_37 | ~ spl55_39)), inference(subsumption_resolution, [], [f6828, f376])).
fof(f6828, plain, (p2(sK44) | p2(sK29(sK44)) | ~ sP0(sK32) | ~ spl55_37), inference(resolution, [], [f368, f147])).
fof(f517, plain, (~ spl55_38 | ~ spl55_40 | ~ spl55_41), inference(avatar_contradiction_clause, [], [f516])).
fof(f516, plain, ($false | (~ spl55_38 | ~ spl55_40 | ~ spl55_41)), inference(subsumption_resolution, [], [f515, f497])).
fof(f497, plain, (p5(sK48(sK32)) | ~ spl55_38), inference(resolution, [], [f372, f94])).
fof(f94, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL660+1.005.p', reflexivity)).
fof(f372, plain, (! [X35] : (~ r1(sK32, X35) | p5(sK48(X35))) | ~ spl55_38), inference(avatar_component_clause, [], [f371])).
fof(f371, plain, (spl55_38 <=> ! [X35] : (p5(sK48(X35)) | ~ r1(sK32, X35))), introduced(avatar_definition, [new_symbols(naming, [spl55_38])])).
fof(f515, plain, (~ p5(sK48(sK32)) | (~ spl55_40 | ~ spl55_41)), inference(subsumption_resolution, [], [f503, f94])).
fof(f503, plain, (~ r1(sK32, sK32) | ~ p5(sK48(sK32)) | (~ spl55_40 | ~ spl55_41)), inference(resolution, [], [f380, f384])).
fof(f384, plain, (! [X45] : (~ r1(sK32, X45) | ~ p5(X45)) | ~ spl55_41), inference(avatar_component_clause, [], [f383])).
fof(f383, plain, (spl55_41 <=> ! [X45] : (~ p5(X45) | ~ r1(sK32, X45))), introduced(avatar_definition, [new_symbols(naming, [spl55_41])])).
fof(f380, plain, (! [X35] : (r1(X35, sK48(X35)) | ~ r1(sK32, X35)) | ~ spl55_40), inference(avatar_component_clause, [], [f379])).
fof(f379, plain, (spl55_40 <=> ! [X35] : (r1(X35, sK48(X35)) | ~ r1(sK32, X35))), introduced(avatar_definition, [new_symbols(naming, [spl55_40])])).
fof(f410, plain, (spl55_41 | spl55_47), inference(avatar_split_clause, [], [f148, f407, f383])).
fof(f148, plain, ! [X45] : (r1(sK32, sK54) | ~ p5(X45) | ~ r1(sK32, X45)), inference(cnf_transformation, [], [f93])).
fof(f405, plain, (spl55_41 | ~ spl55_46), inference(avatar_split_clause, [], [f149, f402, f383])).
fof(f149, plain, ! [X45] : (~ p2(sK54) | ~ p5(X45) | ~ r1(sK32, X45)), inference(cnf_transformation, [], [f93])).
fof(f400, plain, (spl55_41 | spl55_45), inference(avatar_split_clause, [], [f150, f398, f383])).
fof(f150, plain, ! [X45, X41] : (r1(X41, sK52(X41)) | p2(X41) | ~ r1(sK32, X41) | ~ p5(X45) | ~ r1(sK32, X45)), inference(cnf_transformation, [], [f93])).
fof(f396, plain, (spl55_41 | spl55_44), inference(avatar_split_clause, [], [f151, f394, f383])).
fof(f151, plain, ! [X45, X41] : (r1(sK52(X41), sK53(X41)) | p2(X41) | ~ r1(sK32, X41) | ~ p5(X45) | ~ r1(sK32, X45)), inference(cnf_transformation, [], [f93])).
fof(f392, plain, (spl55_41 | spl55_43), inference(avatar_split_clause, [], [f152, f390, f383])).
fof(f152, plain, ! [X45, X41] : (~ p2(sK53(X41)) | p2(X41) | ~ r1(sK32, X41) | ~ p5(X45) | ~ r1(sK32, X45)), inference(cnf_transformation, [], [f93])).
fof(f388, plain, (spl55_41 | spl55_42), inference(avatar_split_clause, [], [f153, f386, f383])).
fof(f153, plain, ! [X45, X41] : (p2(sK52(X41)) | p2(X41) | ~ r1(sK32, X41) | ~ p5(X45) | ~ r1(sK32, X45)), inference(cnf_transformation, [], [f93])).
fof(f381, plain, (spl55_40 | spl55_39), inference(avatar_split_clause, [], [f160, f374, f379])).
fof(f160, plain, ! [X35] : (sP0(sK32) | r1(X35, sK48(X35)) | ~ r1(sK32, X35)), inference(cnf_transformation, [], [f93])).
fof(f377, plain, (spl55_38 | spl55_39), inference(avatar_split_clause, [], [f161, f374, f371])).
fof(f161, plain, ! [X35] : (sP0(sK32) | p5(sK48(X35)) | ~ r1(sK32, X35)), inference(cnf_transformation, [], [f93])).
fof(f369, plain, (spl55_31 | spl55_37), inference(avatar_split_clause, [], [f168, f366, f341])).
fof(f168, plain, (r1(sK32, sK44) | sP6(sK32)), inference(cnf_transformation, [], [f93])).
fof(f364, plain, (spl55_31 | spl55_34 | ~ spl55_36), inference(avatar_split_clause, [], [f169, f361, f353, f341])).
fof(f169, plain, (~ p2(sK44) | sP2(sK44) | sP6(sK32)), inference(cnf_transformation, [], [f93])).
fof(f359, plain, (spl55_31 | spl55_34 | spl55_35), inference(avatar_split_clause, [], [f170, f357, f353, f341])).
fof(f170, plain, ! [X30, X29] : (~ p2(X29) | p2(X30) | ~ r1(X29, X30) | ~ r1(sK44, X29) | sP2(sK44) | sP6(sK32)), inference(cnf_transformation, [], [f93])).
fof(f351, plain, (spl55_31 | spl55_33), inference(avatar_split_clause, [], [f171, f349, f341])).
fof(f171, plain, ! [X26] : (~ p2(X26) | sP4(X26) | sP5(X26) | ~ r1(sK44, X26) | sP6(sK32)), inference(cnf_transformation, [], [f93])).
fof(f347, plain, (spl55_31 | spl55_32), inference(avatar_split_clause, [], [f172, f345, f341])).
fof(f172, plain, ! [X28, X26, X27] : (~ p2(X27) | p2(X28) | ~ r1(X27, X28) | ~ r1(X26, X27) | sP4(X26) | sP5(X26) | ~ r1(sK44, X26) | sP6(sK32)), inference(cnf_transformation, [], [f93])).