fof(f74, plain, $false, inference(subsumption_resolution, [], [f73, f69])).
fof(f69, plain, ~ p2(sK5), inference(subsumption_resolution, [], [f68, f44])).
fof(f44, plain, ~ p101(sK3), inference(cnf_transformation, [], [f19])).
fof(f19, plain, (p100(sK3) & ~ p101(sK3) & (~ p101(sK3) | p100(sK3)) & (~ p100(sK3) | ((~ p1(sK3) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(sK3, X1))) & (p1(sK3) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(sK3, X2))))) & (~ p101(sK3) | ((~ p2(sK3) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(sK3, X3))) & (p2(sK3) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(sK3, X4))))) & (~ p100(sK3) | p101(sK3) | ((p101(sK4) & p2(sK4) & r1(sK3, sK4)) & (p101(sK5) & ~ p2(sK5) & r1(sK3, sK5)))) & ! [X7] : (((~ p101(X7) | p100(X7)) & (~ p100(X7) | ((~ p1(X7) | ! [X8] : (~ p100(X8) | p1(X8) | ~ r1(X7, X8))) & (p1(X7) | ! [X9] : (~ p100(X9) | ~ p1(X9) | ~ r1(X7, X9))))) & (~ p101(X7) | ((~ p2(X7) | ! [X10] : (~ p101(X10) | p2(X10) | ~ r1(X7, X10))) & (p2(X7) | ! [X11] : (~ p101(X11) | ~ p2(X11) | ~ r1(X7, X11))))) & sP0(X7)) | ~ r1(sK3, X7)) & ! [X12] : (p2(X12) | ~ r1(sK3, X12))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3, sK4, sK5])], [f15, f18, f17, f16])).
fof(f16, plain, (? [X0] : (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p100(X0) | p101(X0) | (? [X5] : (p101(X5) & p2(X5) & r1(X0, X5)) & ? [X6] : (p101(X6) & ~ p2(X6) & r1(X0, X6)))) & ! [X7] : (((~ p101(X7) | p100(X7)) & (~ p100(X7) | ((~ p1(X7) | ! [X8] : (~ p100(X8) | p1(X8) | ~ r1(X7, X8))) & (p1(X7) | ! [X9] : (~ p100(X9) | ~ p1(X9) | ~ r1(X7, X9))))) & (~ p101(X7) | ((~ p2(X7) | ! [X10] : (~ p101(X10) | p2(X10) | ~ r1(X7, X10))) & (p2(X7) | ! [X11] : (~ p101(X11) | ~ p2(X11) | ~ r1(X7, X11))))) & sP0(X7)) | ~ r1(X0, X7)) & ! [X12] : (p2(X12) | ~ r1(X0, X12))) => (p100(sK3) & ~ p101(sK3) & (~ p101(sK3) | p100(sK3)) & (~ p100(sK3) | ((~ p1(sK3) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(sK3, X1))) & (p1(sK3) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(sK3, X2))))) & (~ p101(sK3) | ((~ p2(sK3) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(sK3, X3))) & (p2(sK3) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(sK3, X4))))) & (~ p100(sK3) | p101(sK3) | (? [X5] : (p101(X5) & p2(X5) & r1(sK3, X5)) & ? [X6] : (p101(X6) & ~ p2(X6) & r1(sK3, X6)))) & ! [X7] : (((~ p101(X7) | p100(X7)) & (~ p100(X7) | ((~ p1(X7) | ! [X8] : (~ p100(X8) | p1(X8) | ~ r1(X7, X8))) & (p1(X7) | ! [X9] : (~ p100(X9) | ~ p1(X9) | ~ r1(X7, X9))))) & (~ p101(X7) | ((~ p2(X7) | ! [X10] : (~ p101(X10) | p2(X10) | ~ r1(X7, X10))) & (p2(X7) | ! [X11] : (~ p101(X11) | ~ p2(X11) | ~ r1(X7, X11))))) & sP0(X7)) | ~ r1(sK3, X7)) & ! [X12] : (p2(X12) | ~ r1(sK3, X12)))), introduced(choice_axiom, [])).
fof(f17, plain, (? [X5] : (p101(X5) & p2(X5) & r1(sK3, X5)) => (p101(sK4) & p2(sK4) & r1(sK3, sK4))), introduced(choice_axiom, [])).
fof(f18, plain, (? [X6] : (p101(X6) & ~ p2(X6) & r1(sK3, X6)) => (p101(sK5) & ~ p2(sK5) & r1(sK3, sK5))), introduced(choice_axiom, [])).
fof(f15, plain, ? [X0] : (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p100(X0) | p101(X0) | (? [X5] : (p101(X5) & p2(X5) & r1(X0, X5)) & ? [X6] : (p101(X6) & ~ p2(X6) & r1(X0, X6)))) & ! [X7] : (((~ p101(X7) | p100(X7)) & (~ p100(X7) | ((~ p1(X7) | ! [X8] : (~ p100(X8) | p1(X8) | ~ r1(X7, X8))) & (p1(X7) | ! [X9] : (~ p100(X9) | ~ p1(X9) | ~ r1(X7, X9))))) & (~ p101(X7) | ((~ p2(X7) | ! [X10] : (~ p101(X10) | p2(X10) | ~ r1(X7, X10))) & (p2(X7) | ! [X11] : (~ p101(X11) | ~ p2(X11) | ~ r1(X7, X11))))) & sP0(X7)) | ~ r1(X0, X7)) & ! [X12] : (p2(X12) | ~ r1(X0, X12))), inference(rectify, [], [f9])).
fof(f9, plain, ? [X0] : (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p100(X0) | p101(X0) | (? [X5] : (p101(X5) & p2(X5) & r1(X0, X5)) & ? [X6] : (p101(X6) & ~ p2(X6) & r1(X0, X6)))) & ! [X7] : (((~ p101(X7) | p100(X7)) & (~ p100(X7) | ((~ p1(X7) | ! [X8] : (~ p100(X8) | p1(X8) | ~ r1(X7, X8))) & (p1(X7) | ! [X9] : (~ p100(X9) | ~ p1(X9) | ~ r1(X7, X9))))) & (~ p101(X7) | ((~ p2(X7) | ! [X10] : (~ p101(X10) | p2(X10) | ~ r1(X7, X10))) & (p2(X7) | ! [X11] : (~ p101(X11) | ~ p2(X11) | ~ r1(X7, X11))))) & sP0(X7)) | ~ r1(X0, X7)) & ! [X14] : (p2(X14) | ~ r1(X0, X14))), inference(definition_folding, [], [f7, e8])).
fof(f8, plain, ! [X7] : (~ p100(X7) | p101(X7) | (? [X12] : (p101(X12) & p2(X12) & r1(X7, X12)) & ? [X13] : (p101(X13) & ~ p2(X13) & r1(X7, X13))) | ~ sP0(X7)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X7] : (sP0(X7) <=> (~ p100(X7) | p101(X7) | (? [X12] : (p101(X12) & p2(X12) & r1(X7, X12)) & ? [X13] : (p101(X13) & ~ p2(X13) & r1(X7, X13))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f7, plain, ? [X0] : (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p100(X0) | p101(X0) | (? [X5] : (p101(X5) & p2(X5) & r1(X0, X5)) & ? [X6] : (p101(X6) & ~ p2(X6) & r1(X0, X6)))) & ! [X7] : (((~ p101(X7) | p100(X7)) & (~ p100(X7) | ((~ p1(X7) | ! [X8] : (~ p100(X8) | p1(X8) | ~ r1(X7, X8))) & (p1(X7) | ! [X9] : (~ p100(X9) | ~ p1(X9) | ~ r1(X7, X9))))) & (~ p101(X7) | ((~ p2(X7) | ! [X10] : (~ p101(X10) | p2(X10) | ~ r1(X7, X10))) & (p2(X7) | ! [X11] : (~ p101(X11) | ~ p2(X11) | ~ r1(X7, X11))))) & (~ p100(X7) | p101(X7) | (? [X12] : (p101(X12) & p2(X12) & r1(X7, X12)) & ? [X13] : (p101(X13) & ~ p2(X13) & r1(X7, X13))))) | ~ r1(X0, X7)) & ! [X14] : (p2(X14) | ~ r1(X0, X14))), inference(flattening, [], [f6])).
fof(f6, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & ((~ p100(X0) | p101(X0)) | (? [X5] : ((p101(X5) & p2(X5)) & r1(X0, X5)) & ? [X6] : ((p101(X6) & ~ p2(X6)) & r1(X0, X6)))) & ! [X7] : (((~ p101(X7) | p100(X7)) & (~ p100(X7) | ((~ p1(X7) | ! [X8] : (~ p100(X8) | p1(X8) | ~ r1(X7, X8))) & (p1(X7) | ! [X9] : (~ p100(X9) | ~ p1(X9) | ~ r1(X7, X9))))) & (~ p101(X7) | ((~ p2(X7) | ! [X10] : (~ p101(X10) | p2(X10) | ~ r1(X7, X10))) & (p2(X7) | ! [X11] : (~ p101(X11) | ~ p2(X11) | ~ r1(X7, X11))))) & ((~ p100(X7) | p101(X7)) | (? [X12] : ((p101(X12) & p2(X12)) & r1(X7, X12)) & ? [X13] : ((p101(X13) & ~ p2(X13)) & r1(X7, X13))))) | ~ r1(X0, X7))) & ! [X14] : (p2(X14) | ~ r1(X0, X14))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X5] : (~ (p101(X5) & p2(X5)) | ~ r1(X0, X5)) & ~ ! [X6] : (~ (p101(X6) & ~ p2(X6)) | ~ r1(X0, X6)))) & ! [X7] : (((~ p101(X7) | p100(X7)) & (~ p100(X7) | ((~ p1(X7) | ! [X8] : (~ p100(X8) | p1(X8) | ~ r1(X7, X8))) & (p1(X7) | ! [X9] : (~ p100(X9) | ~ p1(X9) | ~ r1(X7, X9))))) & (~ p101(X7) | ((~ p2(X7) | ! [X10] : (~ p101(X10) | p2(X10) | ~ r1(X7, X10))) & (p2(X7) | ! [X11] : (~ p101(X11) | ~ p2(X11) | ~ r1(X7, X11))))) & (~ (p100(X7) & ~ p101(X7)) | (~ ! [X12] : (~ (p101(X12) & p2(X12)) | ~ r1(X7, X12)) & ~ ! [X13] : (~ (p101(X13) & ~ p2(X13)) | ~ r1(X7, X13))))) | ~ r1(X0, X7))) | ~ ! [X14] : (p2(X14) | ~ r1(X0, X14))), inference(pure_predicate_removal, [], [f4])).
fof(f4, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X5] : (~ (p101(X5) & ~ p102(X5) & p2(X5)) | ~ r1(X0, X5)) & ~ ! [X6] : (~ (p101(X6) & ~ p102(X6) & ~ p2(X6)) | ~ r1(X0, X6)))) & ! [X7] : (((~ p101(X7) | p100(X7)) & (~ p102(X7) | p101(X7)) & (~ p100(X7) | ((~ p1(X7) | ! [X8] : (~ p100(X8) | p1(X8) | ~ r1(X7, X8))) & (p1(X7) | ! [X9] : (~ p100(X9) | ~ p1(X9) | ~ r1(X7, X9))))) & (~ p101(X7) | ((~ p2(X7) | ! [X10] : (~ p101(X10) | p2(X10) | ~ r1(X7, X10))) & (p2(X7) | ! [X11] : (~ p101(X11) | ~ p2(X11) | ~ r1(X7, X11))))) & (~ (p100(X7) & ~ p101(X7)) | (~ ! [X12] : (~ (p101(X12) & ~ p102(X12) & p2(X12)) | ~ r1(X7, X12)) & ~ ! [X13] : (~ (p101(X13) & ~ p102(X13) & ~ p2(X13)) | ~ r1(X7, X13))))) | ~ r1(X0, X7))) | ~ ! [X14] : (p2(X14) | ~ r1(X0, X14))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X5] : (~ (p101(X5) & ~ p102(X5) & p2(X5)) | ~ r1(X0, X5)) & ~ ! [X6] : (~ (p101(X6) & ~ p102(X6) & ~ p2(X6)) | ~ r1(X0, X6)))) & ! [X7] : (((~ p101(X7) | p100(X7)) & (~ p102(X7) | p101(X7)) & (~ p100(X7) | ((~ p1(X7) | ! [X8] : (~ p100(X8) | p1(X8) | ~ r1(X7, X8))) & (p1(X7) | ! [X9] : (~ p100(X9) | ~ p1(X9) | ~ r1(X7, X9))))) & (~ p101(X7) | ((~ p2(X7) | ! [X10] : (~ p101(X10) | p2(X10) | ~ r1(X7, X10))) & (p2(X7) | ! [X11] : (~ p101(X11) | ~ p2(X11) | ~ r1(X7, X11))))) & (~ (p100(X7) & ~ p101(X7)) | (~ ! [X12] : (~ (p101(X12) & ~ p102(X12) & p2(X12)) | ~ r1(X7, X12)) & ~ ! [X13] : (~ (p101(X13) & ~ p102(X13) & ~ p2(X13)) | ~ r1(X7, X13))))) | ~ r1(X0, X7))) | ~ ! [X14] : (p2(X14) | ~ r1(X0, X14))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X1] : (~ p100(X1) | ~ p1(X1) | ~ r1(X0, X1))))) & (~ p101(X0) | ((~ p2(X0) | ! [X1] : (~ p101(X1) | p2(X1) | ~ r1(X0, X1))) & (p2(X0) | ! [X1] : (~ p101(X1) | ~ p2(X1) | ~ r1(X0, X1))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & p2(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & ~ p2(X1)) | ~ r1(X0, X1)))) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p2(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X1] : (~ p100(X1) | ~ p1(X1) | ~ r1(X0, X1))))) & (~ p101(X0) | ((~ p2(X0) | ! [X1] : (~ p101(X1) | p2(X1) | ~ r1(X0, X1))) & (p2(X0) | ! [X1] : (~ p101(X1) | ~ p2(X1) | ~ r1(X0, X1))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & p2(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & ~ p2(X1)) | ~ r1(X0, X1)))) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p2(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL636+1.001.p', main)).
fof(f68, plain, (p101(sK3) | ~ p2(sK5)), inference(subsumption_resolution, [], [f34, f45])).
fof(f45, plain, p100(sK3), inference(cnf_transformation, [], [f19])).
fof(f34, plain, (~ p100(sK3) | p101(sK3) | ~ p2(sK5)), inference(cnf_transformation, [], [f19])).
fof(f73, plain, p2(sK5), inference(resolution, [], [f26, f71])).
fof(f71, plain, r1(sK3, sK5), inference(subsumption_resolution, [], [f70, f44])).
fof(f70, plain, (p101(sK3) | r1(sK3, sK5)), inference(subsumption_resolution, [], [f33, f45])).
fof(f33, plain, (~ p100(sK3) | p101(sK3) | r1(sK3, sK5)), inference(cnf_transformation, [], [f19])).
fof(f26, plain, ! [X12] : (~ r1(sK3, X12) | p2(X12)), inference(cnf_transformation, [], [f19])).