fof(f367, plain, $false, inference(subsumption_resolution, [], [f366, f143])).
fof(f143, plain, sP6(sK22), inference(resolution, [], [f137, f61])).
fof(f61, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL686+1.005.p', reflexivity)).
fof(f137, plain, ! [X2] : (~ r1(sK22, X2) | sP6(X2)), inference(cnf_transformation, [], [f60])).
fof(f60, plain, ((! [X2] : ((sP5(X2) & (~ p15(sK23(X2)) & r1(X2, sK23(X2))) & ! [X4] : (((~ p1(X4) | p2(X4)) & (p1(X4) | ~ p2(X4))) | ~ r1(X2, X4)) & sP6(X2)) | ~ r1(sK22, X2)) & r1(sK21, sK22)) & ((p1(sK25) & r1(sK24, sK25)) & p15(sK24) & r1(sK21, sK24))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21, sK22, sK23, sK24, sK25])], [f54, f59, f58, f57, f56, f55])).
fof(f55, plain, (? [X0] : (? [X1] : (! [X2] : ((sP5(X2) & ? [X3] : (~ p15(X3) & r1(X2, X3)) & ! [X4] : (((~ p1(X4) | p2(X4)) & (p1(X4) | ~ p2(X4))) | ~ r1(X2, X4)) & sP6(X2)) | ~ r1(X1, X2)) & r1(X0, X1)) & ? [X5] : (? [X6] : (p1(X6) & r1(X5, X6)) & p15(X5) & r1(X0, X5))) => (? [X1] : (! [X2] : ((sP5(X2) & ? [X3] : (~ p15(X3) & r1(X2, X3)) & ! [X4] : (((~ p1(X4) | p2(X4)) & (p1(X4) | ~ p2(X4))) | ~ r1(X2, X4)) & sP6(X2)) | ~ r1(X1, X2)) & r1(sK21, X1)) & ? [X5] : (? [X6] : (p1(X6) & r1(X5, X6)) & p15(X5) & r1(sK21, X5)))), introduced(choice_axiom, [])).
fof(f56, plain, (? [X1] : (! [X2] : ((sP5(X2) & ? [X3] : (~ p15(X3) & r1(X2, X3)) & ! [X4] : (((~ p1(X4) | p2(X4)) & (p1(X4) | ~ p2(X4))) | ~ r1(X2, X4)) & sP6(X2)) | ~ r1(X1, X2)) & r1(sK21, X1)) => (! [X2] : ((sP5(X2) & ? [X3] : (~ p15(X3) & r1(X2, X3)) & ! [X4] : (((~ p1(X4) | p2(X4)) & (p1(X4) | ~ p2(X4))) | ~ r1(X2, X4)) & sP6(X2)) | ~ r1(sK22, X2)) & r1(sK21, sK22))), introduced(choice_axiom, [])).
fof(f57, plain, ! [X2] : (? [X3] : (~ p15(X3) & r1(X2, X3)) => (~ p15(sK23(X2)) & r1(X2, sK23(X2)))), introduced(choice_axiom, [])).
fof(f58, plain, (? [X5] : (? [X6] : (p1(X6) & r1(X5, X6)) & p15(X5) & r1(sK21, X5)) => (? [X6] : (p1(X6) & r1(sK24, X6)) & p15(sK24) & r1(sK21, sK24))), introduced(choice_axiom, [])).
fof(f59, plain, (? [X6] : (p1(X6) & r1(sK24, X6)) => (p1(sK25) & r1(sK24, sK25))), introduced(choice_axiom, [])).
fof(f54, plain, ? [X0] : (? [X1] : (! [X2] : ((sP5(X2) & ? [X3] : (~ p15(X3) & r1(X2, X3)) & ! [X4] : (((~ p1(X4) | p2(X4)) & (p1(X4) | ~ p2(X4))) | ~ r1(X2, X4)) & sP6(X2)) | ~ r1(X1, X2)) & r1(X0, X1)) & ? [X5] : (? [X6] : (p1(X6) & r1(X5, X6)) & p15(X5) & r1(X0, X5))), inference(rectify, [], [f18])).
fof(f18, plain, ? [X0] : (? [X1] : (! [X2] : ((sP5(X2) & ? [X4] : (~ p15(X4) & r1(X2, X4)) & ! [X5] : (((~ p1(X5) | p2(X5)) & (p1(X5) | ~ p2(X5))) | ~ r1(X2, X5)) & sP6(X2)) | ~ r1(X1, X2)) & r1(X0, X1)) & ? [X31] : (? [X32] : (p1(X32) & r1(X31, X32)) & p15(X31) & r1(X0, X31))), inference(definition_folding, [], [f10, e17, e16, e15, e14, e13, e12, e11])).
fof(f11, plain, ! [X24] : (? [X26] : (! [X27] : (((~ p12(X27) | p13(X27)) & (p12(X27) | ~ p13(X27))) | ~ r1(X26, X27)) & ? [X28] : (! [X29] : (((~ p13(X29) | p14(X29)) & (p13(X29) | ~ p14(X29))) | ~ r1(X28, X29)) & ? [X30] : r1(X28, X30) & r1(X26, X28)) & r1(X24, X26)) | ~ sP0(X24)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X24] : (sP0(X24) <=> ? [X26] : (! [X27] : (((~ p12(X27) | p13(X27)) & (p12(X27) | ~ p13(X27))) | ~ r1(X26, X27)) & ? [X28] : (! [X29] : (((~ p13(X29) | p14(X29)) & (p13(X29) | ~ p14(X29))) | ~ r1(X28, X29)) & ? [X30] : r1(X28, X30) & r1(X26, X28)) & r1(X24, X26))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f12, plain, ! [X20] : (? [X22] : (! [X23] : (((~ p10(X23) | p11(X23)) & (p10(X23) | ~ p11(X23))) | ~ r1(X22, X23)) & ? [X24] : (! [X25] : (((~ p11(X25) | p12(X25)) & (p11(X25) | ~ p12(X25))) | ~ r1(X24, X25)) & sP0(X24) & r1(X22, X24)) & r1(X20, X22)) | ~ sP1(X20)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X20] : (sP1(X20) <=> ? [X22] : (! [X23] : (((~ p10(X23) | p11(X23)) & (p10(X23) | ~ p11(X23))) | ~ r1(X22, X23)) & ? [X24] : (! [X25] : (((~ p11(X25) | p12(X25)) & (p11(X25) | ~ p12(X25))) | ~ r1(X24, X25)) & sP0(X24) & r1(X22, X24)) & r1(X20, X22))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f13, plain, ! [X16] : (? [X18] : (! [X19] : (((~ p8(X19) | p9(X19)) & (p8(X19) | ~ p9(X19))) | ~ r1(X18, X19)) & ? [X20] : (! [X21] : (((~ p9(X21) | p10(X21)) & (p9(X21) | ~ p10(X21))) | ~ r1(X20, X21)) & sP1(X20) & r1(X18, X20)) & r1(X16, X18)) | ~ sP2(X16)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X16] : (sP2(X16) <=> ? [X18] : (! [X19] : (((~ p8(X19) | p9(X19)) & (p8(X19) | ~ p9(X19))) | ~ r1(X18, X19)) & ? [X20] : (! [X21] : (((~ p9(X21) | p10(X21)) & (p9(X21) | ~ p10(X21))) | ~ r1(X20, X21)) & sP1(X20) & r1(X18, X20)) & r1(X16, X18))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f14, plain, ! [X12] : (? [X14] : (! [X15] : (((~ p6(X15) | p7(X15)) & (p6(X15) | ~ p7(X15))) | ~ r1(X14, X15)) & ? [X16] : (! [X17] : (((~ p7(X17) | p8(X17)) & (p7(X17) | ~ p8(X17))) | ~ r1(X16, X17)) & sP2(X16) & r1(X14, X16)) & r1(X12, X14)) | ~ sP3(X12)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X12] : (sP3(X12) <=> ? [X14] : (! [X15] : (((~ p6(X15) | p7(X15)) & (p6(X15) | ~ p7(X15))) | ~ r1(X14, X15)) & ? [X16] : (! [X17] : (((~ p7(X17) | p8(X17)) & (p7(X17) | ~ p8(X17))) | ~ r1(X16, X17)) & sP2(X16) & r1(X14, X16)) & r1(X12, X14))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f15, plain, ! [X8] : (? [X10] : (! [X11] : (((~ p4(X11) | p5(X11)) & (p4(X11) | ~ p5(X11))) | ~ r1(X10, X11)) & ? [X12] : (! [X13] : (((~ p5(X13) | p6(X13)) & (p5(X13) | ~ p6(X13))) | ~ r1(X12, X13)) & sP3(X12) & r1(X10, X12)) & r1(X8, X10)) | ~ sP4(X8)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X8] : (sP4(X8) <=> ? [X10] : (! [X11] : (((~ p4(X11) | p5(X11)) & (p4(X11) | ~ p5(X11))) | ~ r1(X10, X11)) & ? [X12] : (! [X13] : (((~ p5(X13) | p6(X13)) & (p5(X13) | ~ p6(X13))) | ~ r1(X12, X13)) & sP3(X12) & r1(X10, X12)) & r1(X8, X10))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f16, plain, ! [X2] : (? [X3] : ((~ p1(X3) | ~ p2(X3)) & (p2(X3) | p1(X3)) & (~ p2(X3) | ~ p3(X3)) & (p3(X3) | p2(X3)) & (~ p3(X3) | ~ p4(X3)) & (p4(X3) | p3(X3)) & (~ p4(X3) | ~ p5(X3)) & (p5(X3) | p4(X3)) & (~ p5(X3) | ~ p6(X3)) & (p6(X3) | p5(X3)) & (~ p6(X3) | ~ p7(X3)) & (p7(X3) | p6(X3)) & (~ p7(X3) | ~ p8(X3)) & (p8(X3) | p7(X3)) & (~ p8(X3) | ~ p9(X3)) & (p9(X3) | p8(X3)) & (~ p9(X3) | ~ p10(X3)) & (p10(X3) | p9(X3)) & (~ p10(X3) | ~ p11(X3)) & (p11(X3) | p10(X3)) & (~ p11(X3) | ~ p12(X3)) & (p12(X3) | p11(X3)) & (~ p12(X3) | ~ p13(X3)) & (p13(X3) | p12(X3)) & (~ p13(X3) | ~ p14(X3)) & (p14(X3) | p13(X3)) & r1(X2, X3)) | ~ sP5(X2)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X2] : (sP5(X2) <=> ? [X3] : ((~ p1(X3) | ~ p2(X3)) & (p2(X3) | p1(X3)) & (~ p2(X3) | ~ p3(X3)) & (p3(X3) | p2(X3)) & (~ p3(X3) | ~ p4(X3)) & (p4(X3) | p3(X3)) & (~ p4(X3) | ~ p5(X3)) & (p5(X3) | p4(X3)) & (~ p5(X3) | ~ p6(X3)) & (p6(X3) | p5(X3)) & (~ p6(X3) | ~ p7(X3)) & (p7(X3) | p6(X3)) & (~ p7(X3) | ~ p8(X3)) & (p8(X3) | p7(X3)) & (~ p8(X3) | ~ p9(X3)) & (p9(X3) | p8(X3)) & (~ p9(X3) | ~ p10(X3)) & (p10(X3) | p9(X3)) & (~ p10(X3) | ~ p11(X3)) & (p11(X3) | p10(X3)) & (~ p11(X3) | ~ p12(X3)) & (p12(X3) | p11(X3)) & (~ p12(X3) | ~ p13(X3)) & (p13(X3) | p12(X3)) & (~ p13(X3) | ~ p14(X3)) & (p14(X3) | p13(X3)) & r1(X2, X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f17, plain, ! [X2] : (? [X6] : (! [X7] : (((~ p2(X7) | p3(X7)) & (p2(X7) | ~ p3(X7))) | ~ r1(X6, X7)) & ? [X8] : (! [X9] : (((~ p3(X9) | p4(X9)) & (p3(X9) | ~ p4(X9))) | ~ r1(X8, X9)) & sP4(X8) & r1(X6, X8)) & r1(X2, X6)) | ~ sP6(X2)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X2] : (sP6(X2) <=> ? [X6] : (! [X7] : (((~ p2(X7) | p3(X7)) & (p2(X7) | ~ p3(X7))) | ~ r1(X6, X7)) & ? [X8] : (! [X9] : (((~ p3(X9) | p4(X9)) & (p3(X9) | ~ p4(X9))) | ~ r1(X8, X9)) & sP4(X8) & r1(X6, X8)) & r1(X2, X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f10, plain, ? [X0] : (? [X1] : (! [X2] : ((? [X3] : ((~ p1(X3) | ~ p2(X3)) & (p2(X3) | p1(X3)) & (~ p2(X3) | ~ p3(X3)) & (p3(X3) | p2(X3)) & (~ p3(X3) | ~ p4(X3)) & (p4(X3) | p3(X3)) & (~ p4(X3) | ~ p5(X3)) & (p5(X3) | p4(X3)) & (~ p5(X3) | ~ p6(X3)) & (p6(X3) | p5(X3)) & (~ p6(X3) | ~ p7(X3)) & (p7(X3) | p6(X3)) & (~ p7(X3) | ~ p8(X3)) & (p8(X3) | p7(X3)) & (~ p8(X3) | ~ p9(X3)) & (p9(X3) | p8(X3)) & (~ p9(X3) | ~ p10(X3)) & (p10(X3) | p9(X3)) & (~ p10(X3) | ~ p11(X3)) & (p11(X3) | p10(X3)) & (~ p11(X3) | ~ p12(X3)) & (p12(X3) | p11(X3)) & (~ p12(X3) | ~ p13(X3)) & (p13(X3) | p12(X3)) & (~ p13(X3) | ~ p14(X3)) & (p14(X3) | p13(X3)) & r1(X2, X3)) & ? [X4] : (~ p15(X4) & r1(X2, X4)) & ! [X5] : (((~ p1(X5) | p2(X5)) & (p1(X5) | ~ p2(X5))) | ~ r1(X2, X5)) & ? [X6] : (! [X7] : (((~ p2(X7) | p3(X7)) & (p2(X7) | ~ p3(X7))) | ~ r1(X6, X7)) & ? [X8] : (! [X9] : (((~ p3(X9) | p4(X9)) & (p3(X9) | ~ p4(X9))) | ~ r1(X8, X9)) & ? [X10] : (! [X11] : (((~ p4(X11) | p5(X11)) & (p4(X11) | ~ p5(X11))) | ~ r1(X10, X11)) & ? [X12] : (! [X13] : (((~ p5(X13) | p6(X13)) & (p5(X13) | ~ p6(X13))) | ~ r1(X12, X13)) & ? [X14] : (! [X15] : (((~ p6(X15) | p7(X15)) & (p6(X15) | ~ p7(X15))) | ~ r1(X14, X15)) & ? [X16] : (! [X17] : (((~ p7(X17) | p8(X17)) & (p7(X17) | ~ p8(X17))) | ~ r1(X16, X17)) & ? [X18] : (! [X19] : (((~ p8(X19) | p9(X19)) & (p8(X19) | ~ p9(X19))) | ~ r1(X18, X19)) & ? [X20] : (! [X21] : (((~ p9(X21) | p10(X21)) & (p9(X21) | ~ p10(X21))) | ~ r1(X20, X21)) & ? [X22] : (! [X23] : (((~ p10(X23) | p11(X23)) & (p10(X23) | ~ p11(X23))) | ~ r1(X22, X23)) & ? [X24] : (! [X25] : (((~ p11(X25) | p12(X25)) & (p11(X25) | ~ p12(X25))) | ~ r1(X24, X25)) & ? [X26] : (! [X27] : (((~ p12(X27) | p13(X27)) & (p12(X27) | ~ p13(X27))) | ~ r1(X26, X27)) & ? [X28] : (! [X29] : (((~ p13(X29) | p14(X29)) & (p13(X29) | ~ p14(X29))) | ~ r1(X28, X29)) & ? [X30] : r1(X28, X30) & r1(X26, X28)) & r1(X24, X26)) & r1(X22, X24)) & r1(X20, X22)) & r1(X18, X20)) & r1(X16, X18)) & r1(X14, X16)) & r1(X12, X14)) & r1(X10, X12)) & r1(X8, X10)) & r1(X6, X8)) & r1(X2, X6))) | ~ r1(X1, X2)) & r1(X0, X1)) & ? [X31] : (? [X32] : (p1(X32) & r1(X31, X32)) & p15(X31) & r1(X0, X31))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ? [X0] : ~ (! [X1] : (~ ! [X2] : (~ (! [X3] : ((p1(X3) & p2(X3)) | (~ p2(X3) & ~ p1(X3)) | (p2(X3) & p3(X3)) | (~ p3(X3) & ~ p2(X3)) | (p3(X3) & p4(X3)) | (~ p4(X3) & ~ p3(X3)) | (p4(X3) & p5(X3)) | (~ p5(X3) & ~ p4(X3)) | (p5(X3) & p6(X3)) | (~ p6(X3) & ~ p5(X3)) | (p6(X3) & p7(X3)) | (~ p7(X3) & ~ p6(X3)) | (p7(X3) & p8(X3)) | (~ p8(X3) & ~ p7(X3)) | (p8(X3) & p9(X3)) | (~ p9(X3) & ~ p8(X3)) | (p9(X3) & p10(X3)) | (~ p10(X3) & ~ p9(X3)) | (p10(X3) & p11(X3)) | (~ p11(X3) & ~ p10(X3)) | (p11(X3) & p12(X3)) | (~ p12(X3) & ~ p11(X3)) | (p12(X3) & p13(X3)) | (~ p13(X3) & ~ p12(X3)) | (p13(X3) & p14(X3)) | (~ p14(X3) & ~ p13(X3)) | ~ r1(X2, X3)) | ! [X4] : (p15(X4) | ~ r1(X2, X4)) | ~ ! [X5] : (~ ((p1(X5) & ~ p2(X5)) | (~ p1(X5) & p2(X5))) | ~ r1(X2, X5)) | ! [X6] : (~ ! [X7] : (~ ((p2(X7) & ~ p3(X7)) | (~ p2(X7) & p3(X7))) | ~ r1(X6, X7)) | ! [X8] : (~ ! [X9] : (~ ((p3(X9) & ~ p4(X9)) | (~ p3(X9) & p4(X9))) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (~ ((p4(X11) & ~ p5(X11)) | (~ p4(X11) & p5(X11))) | ~ r1(X10, X11)) | ! [X12] : (~ ! [X13] : (~ ((p5(X13) & ~ p6(X13)) | (~ p5(X13) & p6(X13))) | ~ r1(X12, X13)) | ! [X14] : (~ ! [X15] : (~ ((p6(X15) & ~ p7(X15)) | (~ p6(X15) & p7(X15))) | ~ r1(X14, X15)) | ! [X16] : (~ ! [X17] : (~ ((p7(X17) & ~ p8(X17)) | (~ p7(X17) & p8(X17))) | ~ r1(X16, X17)) | ! [X18] : (~ ! [X19] : (~ ((p8(X19) & ~ p9(X19)) | (~ p8(X19) & p9(X19))) | ~ r1(X18, X19)) | ! [X20] : (~ ! [X21] : (~ ((p9(X21) & ~ p10(X21)) | (~ p9(X21) & p10(X21))) | ~ r1(X20, X21)) | ! [X22] : (~ ! [X23] : (~ ((p10(X23) & ~ p11(X23)) | (~ p10(X23) & p11(X23))) | ~ r1(X22, X23)) | ! [X24] : (~ ! [X25] : (~ ((p11(X25) & ~ p12(X25)) | (~ p11(X25) & p12(X25))) | ~ r1(X24, X25)) | ! [X26] : (~ ! [X27] : (~ ((p12(X27) & ~ p13(X27)) | (~ p12(X27) & p13(X27))) | ~ r1(X26, X27)) | ! [X28] : (~ ! [X29] : (~ ((p13(X29) & ~ p14(X29)) | (~ p13(X29) & p14(X29))) | ~ r1(X28, X29)) | ! [X30] : ~ r1(X28, X30) | ~ r1(X26, X28)) | ~ r1(X24, X26)) | ~ r1(X22, X24)) | ~ r1(X20, X22)) | ~ r1(X18, X20)) | ~ r1(X16, X18)) | ~ r1(X14, X16)) | ~ r1(X12, X14)) | ~ r1(X10, X12)) | ~ r1(X8, X10)) | ~ r1(X6, X8)) | ~ r1(X2, X6))) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ! [X31] : (! [X32] : (~ p1(X32) | ~ r1(X31, X32)) | ~ p15(X31) | ~ r1(X0, X31))), inference(flattening, [], [f6])).
fof(f6, plain, ~ ~ ? [X0] : ~ (! [X1] : (~ ! [X2] : (~ (! [X3] : ((p1(X3) & p2(X3)) | (~ p2(X3) & ~ p1(X3)) | (p2(X3) & p3(X3)) | (~ p3(X3) & ~ p2(X3)) | (p3(X3) & p4(X3)) | (~ p4(X3) & ~ p3(X3)) | (p4(X3) & p5(X3)) | (~ p5(X3) & ~ p4(X3)) | (p5(X3) & p6(X3)) | (~ p6(X3) & ~ p5(X3)) | (p6(X3) & p7(X3)) | (~ p7(X3) & ~ p6(X3)) | (p7(X3) & p8(X3)) | (~ p8(X3) & ~ p7(X3)) | (p8(X3) & p9(X3)) | (~ p9(X3) & ~ p8(X3)) | (p9(X3) & p10(X3)) | (~ p10(X3) & ~ p9(X3)) | (p10(X3) & p11(X3)) | (~ p11(X3) & ~ p10(X3)) | (p11(X3) & p12(X3)) | (~ p12(X3) & ~ p11(X3)) | (p12(X3) & p13(X3)) | (~ p13(X3) & ~ p12(X3)) | (p13(X3) & p14(X3)) | (~ p14(X3) & ~ p13(X3)) | ~ r1(X2, X3)) | ! [X4] : (p15(X4) | ~ r1(X2, X4)) | ~ ! [X5] : (~ ((p1(X5) & ~ p2(X5)) | (~ p1(X5) & p2(X5))) | ~ r1(X2, X5)) | ! [X6] : (~ ! [X7] : (~ ((p2(X7) & ~ p3(X7)) | (~ p2(X7) & p3(X7))) | ~ r1(X6, X7)) | ! [X8] : (~ ! [X9] : (~ ((p3(X9) & ~ p4(X9)) | (~ p3(X9) & p4(X9))) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (~ ((p4(X11) & ~ p5(X11)) | (~ p4(X11) & p5(X11))) | ~ r1(X10, X11)) | ! [X12] : (~ ! [X13] : (~ ((p5(X13) & ~ p6(X13)) | (~ p5(X13) & p6(X13))) | ~ r1(X12, X13)) | ! [X14] : (~ ! [X15] : (~ ((p6(X15) & ~ p7(X15)) | (~ p6(X15) & p7(X15))) | ~ r1(X14, X15)) | ! [X16] : (~ ! [X17] : (~ ((p7(X17) & ~ p8(X17)) | (~ p7(X17) & p8(X17))) | ~ r1(X16, X17)) | ! [X18] : (~ ! [X19] : (~ ((p8(X19) & ~ p9(X19)) | (~ p8(X19) & p9(X19))) | ~ r1(X18, X19)) | ! [X20] : (~ ! [X21] : (~ ((p9(X21) & ~ p10(X21)) | (~ p9(X21) & p10(X21))) | ~ r1(X20, X21)) | ! [X22] : (~ ! [X23] : (~ ((p10(X23) & ~ p11(X23)) | (~ p10(X23) & p11(X23))) | ~ r1(X22, X23)) | ! [X24] : (~ ! [X25] : (~ ((p11(X25) & ~ p12(X25)) | (~ p11(X25) & p12(X25))) | ~ r1(X24, X25)) | ! [X26] : (~ ! [X27] : (~ ((p12(X27) & ~ p13(X27)) | (~ p12(X27) & p13(X27))) | ~ r1(X26, X27)) | ! [X28] : (~ ! [X29] : (~ ((p13(X29) & ~ p14(X29)) | (~ p13(X29) & p14(X29))) | ~ r1(X28, X29)) | ! [X30] : ~ r1(X28, X30) | ~ r1(X26, X28)) | ~ r1(X24, X26)) | ~ r1(X22, X24)) | ~ r1(X20, X22)) | ~ r1(X18, X20)) | ~ r1(X16, X18)) | ~ r1(X14, X16)) | ~ r1(X12, X14)) | ~ r1(X10, X12)) | ~ r1(X8, X10)) | ~ r1(X6, X8)) | ~ r1(X2, X6))) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ! [X31] : (! [X32] : (~ p1(X32) | ~ r1(X31, X32)) | ~ p15(X31) | ~ r1(X0, X31))), inference(true_and_false_elimination, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (! [X1] : (~ ! [X2] : (~ (! [X3] : ((p1(X3) & p2(X3)) | (~ p2(X3) & ~ p1(X3)) | (p2(X3) & p3(X3)) | (~ p3(X3) & ~ p2(X3)) | (p3(X3) & p4(X3)) | (~ p4(X3) & ~ p3(X3)) | (p4(X3) & p5(X3)) | (~ p5(X3) & ~ p4(X3)) | (p5(X3) & p6(X3)) | (~ p6(X3) & ~ p5(X3)) | (p6(X3) & p7(X3)) | (~ p7(X3) & ~ p6(X3)) | (p7(X3) & p8(X3)) | (~ p8(X3) & ~ p7(X3)) | (p8(X3) & p9(X3)) | (~ p9(X3) & ~ p8(X3)) | (p9(X3) & p10(X3)) | (~ p10(X3) & ~ p9(X3)) | (p10(X3) & p11(X3)) | (~ p11(X3) & ~ p10(X3)) | (p11(X3) & p12(X3)) | (~ p12(X3) & ~ p11(X3)) | (p12(X3) & p13(X3)) | (~ p13(X3) & ~ p12(X3)) | (p13(X3) & p14(X3)) | (~ p14(X3) & ~ p13(X3)) | ~ r1(X2, X3)) | ! [X4] : (p15(X4) | ~ r1(X2, X4)) | ~ ! [X5] : (~ ((p1(X5) & ~ p2(X5)) | (~ p1(X5) & p2(X5))) | ~ r1(X2, X5)) | ! [X6] : (~ ! [X7] : (~ ((p2(X7) & ~ p3(X7)) | (~ p2(X7) & p3(X7))) | ~ r1(X6, X7)) | ! [X8] : (~ ! [X9] : (~ ((p3(X9) & ~ p4(X9)) | (~ p3(X9) & p4(X9))) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (~ ((p4(X11) & ~ p5(X11)) | (~ p4(X11) & p5(X11))) | ~ r1(X10, X11)) | ! [X12] : (~ ! [X13] : (~ ((p5(X13) & ~ p6(X13)) | (~ p5(X13) & p6(X13))) | ~ r1(X12, X13)) | ! [X14] : (~ ! [X15] : (~ ((p6(X15) & ~ p7(X15)) | (~ p6(X15) & p7(X15))) | ~ r1(X14, X15)) | ! [X16] : (~ ! [X17] : (~ ((p7(X17) & ~ p8(X17)) | (~ p7(X17) & p8(X17))) | ~ r1(X16, X17)) | ! [X18] : (~ ! [X19] : (~ ((p8(X19) & ~ p9(X19)) | (~ p8(X19) & p9(X19))) | ~ r1(X18, X19)) | ! [X20] : (~ ! [X21] : (~ ((p9(X21) & ~ p10(X21)) | (~ p9(X21) & p10(X21))) | ~ r1(X20, X21)) | ! [X22] : (~ ! [X23] : (~ ((p10(X23) & ~ p11(X23)) | (~ p10(X23) & p11(X23))) | ~ r1(X22, X23)) | ! [X24] : (~ ! [X25] : (~ ((p11(X25) & ~ p12(X25)) | (~ p11(X25) & p12(X25))) | ~ r1(X24, X25)) | ! [X26] : (~ ! [X27] : (~ ((p12(X27) & ~ p13(X27)) | (~ p12(X27) & p13(X27))) | ~ r1(X26, X27)) | ! [X28] : (~ ! [X29] : (~ ((p13(X29) & ~ p14(X29)) | (~ p13(X29) & p14(X29))) | ~ r1(X28, X29)) | ! [X30] : ($false | ~ r1(X28, X30)) | ~ r1(X26, X28)) | ~ r1(X24, X26)) | ~ r1(X22, X24)) | ~ r1(X20, X22)) | ~ r1(X18, X20)) | ~ r1(X16, X18)) | ~ r1(X14, X16)) | ~ r1(X12, X14)) | ~ r1(X10, X12)) | ~ r1(X8, X10)) | ~ r1(X6, X8)) | ~ r1(X2, X6))) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ! [X31] : (! [X32] : (~ p1(X32) | ~ r1(X31, X32)) | ~ p15(X31) | ~ r1(X0, X31))), inference(rectify, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (! [X1] : (~ ! [X0] : (~ (! [X1] : ((p1(X1) & p2(X1)) | (~ p2(X1) & ~ p1(X1)) | (p2(X1) & p3(X1)) | (~ p3(X1) & ~ p2(X1)) | (p3(X1) & p4(X1)) | (~ p4(X1) & ~ p3(X1)) | (p4(X1) & p5(X1)) | (~ p5(X1) & ~ p4(X1)) | (p5(X1) & p6(X1)) | (~ p6(X1) & ~ p5(X1)) | (p6(X1) & p7(X1)) | (~ p7(X1) & ~ p6(X1)) | (p7(X1) & p8(X1)) | (~ p8(X1) & ~ p7(X1)) | (p8(X1) & p9(X1)) | (~ p9(X1) & ~ p8(X1)) | (p9(X1) & p10(X1)) | (~ p10(X1) & ~ p9(X1)) | (p10(X1) & p11(X1)) | (~ p11(X1) & ~ p10(X1)) | (p11(X1) & p12(X1)) | (~ p12(X1) & ~ p11(X1)) | (p12(X1) & p13(X1)) | (~ p13(X1) & ~ p12(X1)) | (p13(X1) & p14(X1)) | (~ p14(X1) & ~ p13(X1)) | ~ r1(X0, X1)) | ! [X1] : (p15(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ((p1(X1) & ~ p2(X1)) | (~ p1(X1) & p2(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p2(X0) & ~ p3(X0)) | (~ p2(X0) & p3(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p3(X1) & ~ p4(X1)) | (~ p3(X1) & p4(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p4(X0) & ~ p5(X0)) | (~ p4(X0) & p5(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p5(X1) & ~ p6(X1)) | (~ p5(X1) & p6(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p6(X0) & ~ p7(X0)) | (~ p6(X0) & p7(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p7(X1) & ~ p8(X1)) | (~ p7(X1) & p8(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p8(X0) & ~ p9(X0)) | (~ p8(X0) & p9(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p9(X1) & ~ p10(X1)) | (~ p9(X1) & p10(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p10(X0) & ~ p11(X0)) | (~ p10(X0) & p11(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p11(X1) & ~ p12(X1)) | (~ p11(X1) & p12(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p12(X0) & ~ p13(X0)) | (~ p12(X0) & p13(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p13(X1) & ~ p14(X1)) | (~ p13(X1) & p14(X1))) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ p1(X0) | ~ r1(X1, X0)) | ~ p15(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (! [X1] : (~ ! [X0] : (~ (! [X1] : ((p1(X1) & p2(X1)) | (~ p2(X1) & ~ p1(X1)) | (p2(X1) & p3(X1)) | (~ p3(X1) & ~ p2(X1)) | (p3(X1) & p4(X1)) | (~ p4(X1) & ~ p3(X1)) | (p4(X1) & p5(X1)) | (~ p5(X1) & ~ p4(X1)) | (p5(X1) & p6(X1)) | (~ p6(X1) & ~ p5(X1)) | (p6(X1) & p7(X1)) | (~ p7(X1) & ~ p6(X1)) | (p7(X1) & p8(X1)) | (~ p8(X1) & ~ p7(X1)) | (p8(X1) & p9(X1)) | (~ p9(X1) & ~ p8(X1)) | (p9(X1) & p10(X1)) | (~ p10(X1) & ~ p9(X1)) | (p10(X1) & p11(X1)) | (~ p11(X1) & ~ p10(X1)) | (p11(X1) & p12(X1)) | (~ p12(X1) & ~ p11(X1)) | (p12(X1) & p13(X1)) | (~ p13(X1) & ~ p12(X1)) | (p13(X1) & p14(X1)) | (~ p14(X1) & ~ p13(X1)) | ~ r1(X0, X1)) | ! [X1] : (p15(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ((p1(X1) & ~ p2(X1)) | (~ p1(X1) & p2(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p2(X0) & ~ p3(X0)) | (~ p2(X0) & p3(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p3(X1) & ~ p4(X1)) | (~ p3(X1) & p4(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p4(X0) & ~ p5(X0)) | (~ p4(X0) & p5(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p5(X1) & ~ p6(X1)) | (~ p5(X1) & p6(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p6(X0) & ~ p7(X0)) | (~ p6(X0) & p7(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p7(X1) & ~ p8(X1)) | (~ p7(X1) & p8(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p8(X0) & ~ p9(X0)) | (~ p8(X0) & p9(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p9(X1) & ~ p10(X1)) | (~ p9(X1) & p10(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p10(X0) & ~ p11(X0)) | (~ p10(X0) & p11(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p11(X1) & ~ p12(X1)) | (~ p11(X1) & p12(X1))) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ((p12(X0) & ~ p13(X0)) | (~ p12(X0) & p13(X0))) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ ((p13(X1) & ~ p14(X1)) | (~ p13(X1) & p14(X1))) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ p1(X0) | ~ r1(X1, X0)) | ~ p15(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL686+1.005.p', main)).
fof(f366, plain, ~ sP6(sK22), inference(resolution, [], [f365, f147])).
fof(f147, plain, sP5(sK7(sK22)), inference(subsumption_resolution, [], [f145, f143])).
fof(f145, plain, (~ sP6(sK22) | sP5(sK7(sK22))), inference(resolution, [], [f63, f142])).
fof(f142, plain, ! [X2] : (~ r1(sK22, X2) | sP5(X2)), inference(cnf_transformation, [], [f60])).
fof(f63, plain, ! [X0] : (r1(X0, sK7(X0)) | ~ sP6(X0)), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ! [X0] : ((! [X2] : (((~ p2(X2) | p3(X2)) & (p2(X2) | ~ p3(X2))) | ~ r1(sK7(X0), X2)) & (! [X4] : (((~ p3(X4) | p4(X4)) & (p3(X4) | ~ p4(X4))) | ~ r1(sK8(X0), X4)) & sP4(sK8(X0)) & r1(sK7(X0), sK8(X0))) & r1(X0, sK7(X0))) | ~ sP6(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8])], [f20, f22, f21])).
fof(f21, plain, ! [X0] : (? [X1] : (! [X2] : (((~ p2(X2) | p3(X2)) & (p2(X2) | ~ p3(X2))) | ~ r1(X1, X2)) & ? [X3] : (! [X4] : (((~ p3(X4) | p4(X4)) & (p3(X4) | ~ p4(X4))) | ~ r1(X3, X4)) & sP4(X3) & r1(X1, X3)) & r1(X0, X1)) => (! [X2] : (((~ p2(X2) | p3(X2)) & (p2(X2) | ~ p3(X2))) | ~ r1(sK7(X0), X2)) & ? [X3] : (! [X4] : (((~ p3(X4) | p4(X4)) & (p3(X4) | ~ p4(X4))) | ~ r1(X3, X4)) & sP4(X3) & r1(sK7(X0), X3)) & r1(X0, sK7(X0)))), introduced(choice_axiom, [])).
fof(f22, plain, ! [X0] : (? [X3] : (! [X4] : (((~ p3(X4) | p4(X4)) & (p3(X4) | ~ p4(X4))) | ~ r1(X3, X4)) & sP4(X3) & r1(sK7(X0), X3)) => (! [X4] : (((~ p3(X4) | p4(X4)) & (p3(X4) | ~ p4(X4))) | ~ r1(sK8(X0), X4)) & sP4(sK8(X0)) & r1(sK7(X0), sK8(X0)))), introduced(choice_axiom, [])).
fof(f20, plain, ! [X0] : (? [X1] : (! [X2] : (((~ p2(X2) | p3(X2)) & (p2(X2) | ~ p3(X2))) | ~ r1(X1, X2)) & ? [X3] : (! [X4] : (((~ p3(X4) | p4(X4)) & (p3(X4) | ~ p4(X4))) | ~ r1(X3, X4)) & sP4(X3) & r1(X1, X3)) & r1(X0, X1)) | ~ sP6(X0)), inference(rectify, [], [f19])).
fof(f19, plain, ! [X2] : (? [X6] : (! [X7] : (((~ p2(X7) | p3(X7)) & (p2(X7) | ~ p3(X7))) | ~ r1(X6, X7)) & ? [X8] : (! [X9] : (((~ p3(X9) | p4(X9)) & (p3(X9) | ~ p4(X9))) | ~ r1(X8, X9)) & sP4(X8) & r1(X6, X8)) & r1(X2, X6)) | ~ sP6(X2)), inference(nnf_transformation, [], [f17])).
fof(f365, plain, ! [X3] : (~ sP5(sK7(X3)) | ~ sP6(X3)), inference(subsumption_resolution, [], [f364, f351])).
fof(f351, plain, ! [X3] : (~ p3(sK9(sK7(X3))) | ~ sP6(X3) | ~ sP5(sK7(X3))), inference(subsumption_resolution, [], [f343, f94])).
fof(f94, plain, ! [X0] : (~ p2(sK9(X0)) | ~ p3(sK9(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ! [X0] : (((~ p1(sK9(X0)) | ~ p2(sK9(X0))) & (p2(sK9(X0)) | p1(sK9(X0))) & (~ p2(sK9(X0)) | ~ p3(sK9(X0))) & (p3(sK9(X0)) | p2(sK9(X0))) & (~ p3(sK9(X0)) | ~ p4(sK9(X0))) & (p4(sK9(X0)) | p3(sK9(X0))) & (~ p4(sK9(X0)) | ~ p5(sK9(X0))) & (p5(sK9(X0)) | p4(sK9(X0))) & (~ p5(sK9(X0)) | ~ p6(sK9(X0))) & (p6(sK9(X0)) | p5(sK9(X0))) & (~ p6(sK9(X0)) | ~ p7(sK9(X0))) & (p7(sK9(X0)) | p6(sK9(X0))) & (~ p7(sK9(X0)) | ~ p8(sK9(X0))) & (p8(sK9(X0)) | p7(sK9(X0))) & (~ p8(sK9(X0)) | ~ p9(sK9(X0))) & (p9(sK9(X0)) | p8(sK9(X0))) & (~ p9(sK9(X0)) | ~ p10(sK9(X0))) & (p10(sK9(X0)) | p9(sK9(X0))) & (~ p10(sK9(X0)) | ~ p11(sK9(X0))) & (p11(sK9(X0)) | p10(sK9(X0))) & (~ p11(sK9(X0)) | ~ p12(sK9(X0))) & (p12(sK9(X0)) | p11(sK9(X0))) & (~ p12(sK9(X0)) | ~ p13(sK9(X0))) & (p13(sK9(X0)) | p12(sK9(X0))) & (~ p13(sK9(X0)) | ~ p14(sK9(X0))) & (p14(sK9(X0)) | p13(sK9(X0))) & r1(X0, sK9(X0))) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9])], [f25, f26])).
fof(f26, plain, ! [X0] : (? [X1] : ((~ p1(X1) | ~ p2(X1)) & (p2(X1) | p1(X1)) & (~ p2(X1) | ~ p3(X1)) & (p3(X1) | p2(X1)) & (~ p3(X1) | ~ p4(X1)) & (p4(X1) | p3(X1)) & (~ p4(X1) | ~ p5(X1)) & (p5(X1) | p4(X1)) & (~ p5(X1) | ~ p6(X1)) & (p6(X1) | p5(X1)) & (~ p6(X1) | ~ p7(X1)) & (p7(X1) | p6(X1)) & (~ p7(X1) | ~ p8(X1)) & (p8(X1) | p7(X1)) & (~ p8(X1) | ~ p9(X1)) & (p9(X1) | p8(X1)) & (~ p9(X1) | ~ p10(X1)) & (p10(X1) | p9(X1)) & (~ p10(X1) | ~ p11(X1)) & (p11(X1) | p10(X1)) & (~ p11(X1) | ~ p12(X1)) & (p12(X1) | p11(X1)) & (~ p12(X1) | ~ p13(X1)) & (p13(X1) | p12(X1)) & (~ p13(X1) | ~ p14(X1)) & (p14(X1) | p13(X1)) & r1(X0, X1)) => ((~ p1(sK9(X0)) | ~ p2(sK9(X0))) & (p2(sK9(X0)) | p1(sK9(X0))) & (~ p2(sK9(X0)) | ~ p3(sK9(X0))) & (p3(sK9(X0)) | p2(sK9(X0))) & (~ p3(sK9(X0)) | ~ p4(sK9(X0))) & (p4(sK9(X0)) | p3(sK9(X0))) & (~ p4(sK9(X0)) | ~ p5(sK9(X0))) & (p5(sK9(X0)) | p4(sK9(X0))) & (~ p5(sK9(X0)) | ~ p6(sK9(X0))) & (p6(sK9(X0)) | p5(sK9(X0))) & (~ p6(sK9(X0)) | ~ p7(sK9(X0))) & (p7(sK9(X0)) | p6(sK9(X0))) & (~ p7(sK9(X0)) | ~ p8(sK9(X0))) & (p8(sK9(X0)) | p7(sK9(X0))) & (~ p8(sK9(X0)) | ~ p9(sK9(X0))) & (p9(sK9(X0)) | p8(sK9(X0))) & (~ p9(sK9(X0)) | ~ p10(sK9(X0))) & (p10(sK9(X0)) | p9(sK9(X0))) & (~ p10(sK9(X0)) | ~ p11(sK9(X0))) & (p11(sK9(X0)) | p10(sK9(X0))) & (~ p11(sK9(X0)) | ~ p12(sK9(X0))) & (p12(sK9(X0)) | p11(sK9(X0))) & (~ p12(sK9(X0)) | ~ p13(sK9(X0))) & (p13(sK9(X0)) | p12(sK9(X0))) & (~ p13(sK9(X0)) | ~ p14(sK9(X0))) & (p14(sK9(X0)) | p13(sK9(X0))) & r1(X0, sK9(X0)))), introduced(choice_axiom, [])).
fof(f25, plain, ! [X0] : (? [X1] : ((~ p1(X1) | ~ p2(X1)) & (p2(X1) | p1(X1)) & (~ p2(X1) | ~ p3(X1)) & (p3(X1) | p2(X1)) & (~ p3(X1) | ~ p4(X1)) & (p4(X1) | p3(X1)) & (~ p4(X1) | ~ p5(X1)) & (p5(X1) | p4(X1)) & (~ p5(X1) | ~ p6(X1)) & (p6(X1) | p5(X1)) & (~ p6(X1) | ~ p7(X1)) & (p7(X1) | p6(X1)) & (~ p7(X1) | ~ p8(X1)) & (p8(X1) | p7(X1)) & (~ p8(X1) | ~ p9(X1)) & (p9(X1) | p8(X1)) & (~ p9(X1) | ~ p10(X1)) & (p10(X1) | p9(X1)) & (~ p10(X1) | ~ p11(X1)) & (p11(X1) | p10(X1)) & (~ p11(X1) | ~ p12(X1)) & (p12(X1) | p11(X1)) & (~ p12(X1) | ~ p13(X1)) & (p13(X1) | p12(X1)) & (~ p13(X1) | ~ p14(X1)) & (p14(X1) | p13(X1)) & r1(X0, X1)) | ~ sP5(X0)), inference(rectify, [], [f24])).
fof(f24, plain, ! [X2] : (? [X3] : ((~ p1(X3) | ~ p2(X3)) & (p2(X3) | p1(X3)) & (~ p2(X3) | ~ p3(X3)) & (p3(X3) | p2(X3)) & (~ p3(X3) | ~ p4(X3)) & (p4(X3) | p3(X3)) & (~ p4(X3) | ~ p5(X3)) & (p5(X3) | p4(X3)) & (~ p5(X3) | ~ p6(X3)) & (p6(X3) | p5(X3)) & (~ p6(X3) | ~ p7(X3)) & (p7(X3) | p6(X3)) & (~ p7(X3) | ~ p8(X3)) & (p8(X3) | p7(X3)) & (~ p8(X3) | ~ p9(X3)) & (p9(X3) | p8(X3)) & (~ p9(X3) | ~ p10(X3)) & (p10(X3) | p9(X3)) & (~ p10(X3) | ~ p11(X3)) & (p11(X3) | p10(X3)) & (~ p11(X3) | ~ p12(X3)) & (p12(X3) | p11(X3)) & (~ p12(X3) | ~ p13(X3)) & (p13(X3) | p12(X3)) & (~ p13(X3) | ~ p14(X3)) & (p14(X3) | p13(X3)) & r1(X2, X3)) | ~ sP5(X2)), inference(nnf_transformation, [], [f16])).
fof(f343, plain, ! [X3] : (~ p3(sK9(sK7(X3))) | p2(sK9(sK7(X3))) | ~ sP6(X3) | ~ sP5(sK7(X3))), inference(resolution, [], [f68, f70])).
fof(f70, plain, ! [X0] : (r1(X0, sK9(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f27])).
fof(f68, plain, ! [X2, X0] : (~ r1(sK7(X0), X2) | ~ p3(X2) | p2(X2) | ~ sP6(X0)), inference(cnf_transformation, [], [f23])).
fof(f364, plain, ! [X3] : (p3(sK9(sK7(X3))) | ~ sP6(X3) | ~ sP5(sK7(X3))), inference(subsumption_resolution, [], [f356, f93])).
fof(f93, plain, ! [X0] : (p2(sK9(X0)) | p3(sK9(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f27])).
fof(f356, plain, ! [X3] : (p3(sK9(sK7(X3))) | ~ p2(sK9(sK7(X3))) | ~ sP6(X3) | ~ sP5(sK7(X3))), inference(resolution, [], [f69, f70])).
fof(f69, plain, ! [X2, X0] : (~ r1(sK7(X0), X2) | p3(X2) | ~ p2(X2) | ~ sP6(X0)), inference(cnf_transformation, [], [f23])).