fof(f201172, plain, $false, inference(avatar_sat_refutation, [], [f391, f395, f403, f408, f418, f2914, f3031, f3033, f11685, f31557, f44635, f56636, f97999, f141970, f176976, f177473, f177478, f177689, f177712, f177740, f185469, f201161, f201167, f201171])).
fof(f201171, plain, (~ spl58_171 | spl58_771 | ~ spl58_4715), inference(avatar_contradiction_clause, [], [f201170])).
fof(f201170, plain, ($false | (~ spl58_171 | spl58_771 | ~ spl58_4715)), inference(subsumption_resolution, [], [f201169, f2894])).
fof(f2894, plain, (r1(sK36, sK32(sK36)) | ~ spl58_171), inference(avatar_component_clause, [], [f2893])).
fof(f2893, plain, (spl58_171 <=> r1(sK36, sK32(sK36))), introduced(avatar_definition, [new_symbols(naming, [spl58_171])])).
fof(f201169, plain, (~ r1(sK36, sK32(sK36)) | (spl58_771 | ~ spl58_4715)), inference(subsumption_resolution, [], [f201168, f23149])).
fof(f23149, plain, (~ p2(sK32(sK36)) | spl58_771), inference(avatar_component_clause, [], [f23148])).
fof(f23148, plain, (spl58_771 <=> p2(sK32(sK36))), introduced(avatar_definition, [new_symbols(naming, [spl58_771])])).
fof(f201168, plain, (p2(sK32(sK36)) | ~ r1(sK36, sK32(sK36)) | ~ spl58_4715), inference(resolution, [], [f177765, f213])).
fof(f213, plain, ! [X1] : (~ p2(sK37(X1)) | p2(X1) | ~ r1(sK36, X1)), inference(cnf_transformation, [], [f102])).
fof(f102, plain, (! [X1] : ((! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK37(X1), X3)) & ~ p2(sK37(X1)) & r1(X1, sK37(X1))) | p2(X1) | ~ r1(sK36, X1)) & ((! [X6] : ((r1(X6, sK39(X6)) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(sK38, X6)) & r1(sK38, sK40) & ~ p1(sK38) & r1(sK36, sK38)) | ! [X11] : ~ r1(sK36, X11) | p1(sK36)) & ((! [X13] : ((r1(X13, sK42(X13)) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(sK41, X13)) & r1(sK41, sK43) & ~ p1(sK41) & ~ p2(sK41) & r1(sK36, sK41)) | ! [X18] : ~ r1(sK36, X18) | p1(sK36) | p2(sK36)) & ((sP9(sK44) & r1(sK44, sK45) & ~ p1(sK44) & ~ p2(sK44) & ~ p3(sK44) & r1(sK36, sK44)) | ! [X21] : ~ r1(sK36, X21) | p1(sK36) | p2(sK36) | p3(sK36)) & ((sP8(sK46) & r1(sK46, sK47) & ~ p1(sK46) & ~ p2(sK46) & ~ p3(sK46) & ~ p4(sK46) & r1(sK36, sK46)) | ! [X24] : ~ r1(sK36, X24) | p1(sK36) | p2(sK36) | p3(sK36) | p4(sK36)) & ((! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP5(X26) | sP6(X26) | ~ r1(sK48, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(sK48, X29)) & ~ p2(sK48)) | sP3(sK48)) & r1(sK36, sK48)) | sP7(sK36)) & ! [X31] : ((p1(sK49(X31)) & (~ p1(sK50(X31)) & r1(sK49(X31), sK50(X31))) & r1(X31, sK49(X31))) | p1(X31) | ~ r1(sK36, X31)) & (~ p1(sK51) & r1(sK36, sK51)) & (sP1(sK36) | ! [X35] : ((p5(sK52(X35)) & r1(X35, sK52(X35))) | ~ r1(sK36, X35))) & ! [X37] : ((p3(sK53(X37)) & (~ p3(sK54(X37)) & r1(sK53(X37), sK54(X37))) & r1(X37, sK53(X37))) | p3(X37) | ~ r1(sK36, X37)) & (~ p3(sK55) & r1(sK36, sK55)) & (sP0(sK36) | ((! [X43] : (~ p5(X43) | ~ r1(sK57, X43)) & r1(sK56, sK57)) & ~ p1(sK56) & r1(sK36, sK56)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK36, sK37, sK38, sK39, sK40, sK41, sK42, sK43, sK44, sK45, sK46, sK47, sK48, sK49, sK50, sK51, sK52, sK53, sK54, sK55, sK56, sK57])], [f79, f101, f100, f99, f98, f97, f96, f95, f94, f93, f92, f91, f90, f89, f88, f87, f86, f85, f84, f83, f82, f81, f80])).
fof(f80, plain, (? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP9(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X21] : ~ r1(X0, X21) | p1(X0) | p2(X0) | p3(X0)) & (? [X22] : (sP8(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(X0, X22)) | ! [X24] : ~ r1(X0, X24) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP5(X26) | sP6(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP3(X25)) & r1(X0, X25)) | sP7(X0)) & ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) | p1(X31) | ~ r1(X0, X31)) & ? [X34] : (~ p1(X34) & r1(X0, X34)) & (sP1(X0) | ! [X35] : (? [X36] : (p5(X36) & r1(X35, X36)) | ~ r1(X0, X35))) & ! [X37] : (? [X38] : (p3(X38) & ? [X39] : (~ p3(X39) & r1(X38, X39)) & r1(X37, X38)) | p3(X37) | ~ r1(X0, X37)) & ? [X40] : (~ p3(X40) & r1(X0, X40)) & (sP0(X0) | ? [X41] : (? [X42] : (! [X43] : (~ p5(X43) | ~ r1(X42, X43)) & r1(X41, X42)) & ~ p1(X41) & r1(X0, X41)))) => (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(sK36, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(sK36, X5)) | ! [X11] : ~ r1(sK36, X11) | p1(sK36)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(sK36, X12)) | ! [X18] : ~ r1(sK36, X18) | p1(sK36) | p2(sK36)) & (? [X19] : (sP9(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(sK36, X19)) | ! [X21] : ~ r1(sK36, X21) | p1(sK36) | p2(sK36) | p3(sK36)) & (? [X22] : (sP8(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(sK36, X22)) | ! [X24] : ~ r1(sK36, X24) | p1(sK36) | p2(sK36) | p3(sK36) | p4(sK36)) & (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP5(X26) | sP6(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP3(X25)) & r1(sK36, X25)) | sP7(sK36)) & ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) | p1(X31) | ~ r1(sK36, X31)) & ? [X34] : (~ p1(X34) & r1(sK36, X34)) & (sP1(sK36) | ! [X35] : (? [X36] : (p5(X36) & r1(X35, X36)) | ~ r1(sK36, X35))) & ! [X37] : (? [X38] : (p3(X38) & ? [X39] : (~ p3(X39) & r1(X38, X39)) & r1(X37, X38)) | p3(X37) | ~ r1(sK36, X37)) & ? [X40] : (~ p3(X40) & r1(sK36, X40)) & (sP0(sK36) | ? [X41] : (? [X42] : (! [X43] : (~ p5(X43) | ~ r1(X42, X43)) & r1(X41, X42)) & ~ p1(X41) & r1(sK36, X41))))), introduced(choice_axiom, [])).
fof(f81, plain, ! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) => (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK37(X1), X3)) & ~ p2(sK37(X1)) & r1(X1, sK37(X1)))), introduced(choice_axiom, [])).
fof(f82, plain, (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(sK36, X5)) => (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(sK38, X6)) & ? [X10] : r1(sK38, X10) & ~ p1(sK38) & r1(sK36, sK38))), introduced(choice_axiom, [])).
fof(f83, plain, ! [X6] : (? [X7] : r1(X6, X7) => r1(X6, sK39(X6))), introduced(choice_axiom, [])).
fof(f84, plain, (? [X10] : r1(sK38, X10) => r1(sK38, sK40)), introduced(choice_axiom, [])).
fof(f85, plain, (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(sK36, X12)) => (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(sK41, X13)) & ? [X17] : r1(sK41, X17) & ~ p1(sK41) & ~ p2(sK41) & r1(sK36, sK41))), introduced(choice_axiom, [])).
fof(f86, plain, ! [X13] : (? [X14] : r1(X13, X14) => r1(X13, sK42(X13))), introduced(choice_axiom, [])).
fof(f87, plain, (? [X17] : r1(sK41, X17) => r1(sK41, sK43)), introduced(choice_axiom, [])).
fof(f88, plain, (? [X19] : (sP9(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(sK36, X19)) => (sP9(sK44) & ? [X20] : r1(sK44, X20) & ~ p1(sK44) & ~ p2(sK44) & ~ p3(sK44) & r1(sK36, sK44))), introduced(choice_axiom, [])).
fof(f89, plain, (? [X20] : r1(sK44, X20) => r1(sK44, sK45)), introduced(choice_axiom, [])).
fof(f90, plain, (? [X22] : (sP8(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(sK36, X22)) => (sP8(sK46) & ? [X23] : r1(sK46, X23) & ~ p1(sK46) & ~ p2(sK46) & ~ p3(sK46) & ~ p4(sK46) & r1(sK36, sK46))), introduced(choice_axiom, [])).
fof(f91, plain, (? [X23] : r1(sK46, X23) => r1(sK46, sK47)), introduced(choice_axiom, [])).
fof(f92, plain, (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP5(X26) | sP6(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP3(X25)) & r1(sK36, X25)) => (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP5(X26) | sP6(X26) | ~ r1(sK48, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(sK48, X29)) & ~ p2(sK48)) | sP3(sK48)) & r1(sK36, sK48))), introduced(choice_axiom, [])).
fof(f93, plain, ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) => (p1(sK49(X31)) & ? [X33] : (~ p1(X33) & r1(sK49(X31), X33)) & r1(X31, sK49(X31)))), introduced(choice_axiom, [])).
fof(f94, plain, ! [X31] : (? [X33] : (~ p1(X33) & r1(sK49(X31), X33)) => (~ p1(sK50(X31)) & r1(sK49(X31), sK50(X31)))), introduced(choice_axiom, [])).
fof(f95, plain, (? [X34] : (~ p1(X34) & r1(sK36, X34)) => (~ p1(sK51) & r1(sK36, sK51))), introduced(choice_axiom, [])).
fof(f96, plain, ! [X35] : (? [X36] : (p5(X36) & r1(X35, X36)) => (p5(sK52(X35)) & r1(X35, sK52(X35)))), introduced(choice_axiom, [])).
fof(f97, plain, ! [X37] : (? [X38] : (p3(X38) & ? [X39] : (~ p3(X39) & r1(X38, X39)) & r1(X37, X38)) => (p3(sK53(X37)) & ? [X39] : (~ p3(X39) & r1(sK53(X37), X39)) & r1(X37, sK53(X37)))), introduced(choice_axiom, [])).
fof(f98, plain, ! [X37] : (? [X39] : (~ p3(X39) & r1(sK53(X37), X39)) => (~ p3(sK54(X37)) & r1(sK53(X37), sK54(X37)))), introduced(choice_axiom, [])).
fof(f99, plain, (? [X40] : (~ p3(X40) & r1(sK36, X40)) => (~ p3(sK55) & r1(sK36, sK55))), introduced(choice_axiom, [])).
fof(f100, plain, (? [X41] : (? [X42] : (! [X43] : (~ p5(X43) | ~ r1(X42, X43)) & r1(X41, X42)) & ~ p1(X41) & r1(sK36, X41)) => (? [X42] : (! [X43] : (~ p5(X43) | ~ r1(X42, X43)) & r1(sK56, X42)) & ~ p1(sK56) & r1(sK36, sK56))), introduced(choice_axiom, [])).
fof(f101, plain, (? [X42] : (! [X43] : (~ p5(X43) | ~ r1(X42, X43)) & r1(sK56, X42)) => (! [X43] : (~ p5(X43) | ~ r1(sK57, X43)) & r1(sK56, sK57))), introduced(choice_axiom, [])).
fof(f79, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP9(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X21] : ~ r1(X0, X21) | p1(X0) | p2(X0) | p3(X0)) & (? [X22] : (sP8(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(X0, X22)) | ! [X24] : ~ r1(X0, X24) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP5(X26) | sP6(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP3(X25)) & r1(X0, X25)) | sP7(X0)) & ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) | p1(X31) | ~ r1(X0, X31)) & ? [X34] : (~ p1(X34) & r1(X0, X34)) & (sP1(X0) | ! [X35] : (? [X36] : (p5(X36) & r1(X35, X36)) | ~ r1(X0, X35))) & ! [X37] : (? [X38] : (p3(X38) & ? [X39] : (~ p3(X39) & r1(X38, X39)) & r1(X37, X38)) | p3(X37) | ~ r1(X0, X37)) & ? [X40] : (~ p3(X40) & r1(X0, X40)) & (sP0(X0) | ? [X41] : (? [X42] : (! [X43] : (~ p5(X43) | ~ r1(X42, X43)) & r1(X41, X42)) & ~ p1(X41) & r1(X0, X41)))), inference(rectify, [], [f22])).
fof(f22, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP9(X19) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (sP8(X26) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : ((! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p2(X34)) | sP5(X34) | sP6(X34) | ~ r1(X33, X34)) & ((! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) & ~ p2(X33)) | sP3(X33)) & r1(X0, X33)) | sP7(X0)) & ! [X72] : (? [X73] : (p1(X73) & ? [X74] : (~ p1(X74) & r1(X73, X74)) & r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) & ? [X75] : (~ p1(X75) & r1(X0, X75)) & (sP1(X0) | ! [X80] : (? [X81] : (p5(X81) & r1(X80, X81)) | ~ r1(X0, X80))) & ! [X82] : (? [X83] : (p3(X83) & ? [X84] : (~ p3(X84) & r1(X83, X84)) & r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) & ? [X85] : (~ p3(X85) & r1(X0, X85)) & (sP0(X0) | ? [X90] : (? [X91] : (! [X92] : (~ p5(X92) | ~ r1(X91, X92)) & r1(X90, X91)) & ~ p1(X90) & r1(X0, X90)))), inference(definition_folding, [], [f11, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12])).
fof(f12, plain, ! [X0] : ((! [X86] : (? [X87] : (p2(X87) & ? [X88] : (~ p2(X88) & r1(X87, X88)) & r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) & ? [X89] : (~ p2(X89) & r1(X0, X89))) | ~ sP0(X0)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X0] : (sP0(X0) <=> (! [X86] : (? [X87] : (p2(X87) & ? [X88] : (~ p2(X88) & r1(X87, X88)) & r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) & ? [X89] : (~ p2(X89) & r1(X0, X89)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f13, plain, ! [X0] : ((! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79))) | ~ sP1(X0)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X0] : (sP1(X0) <=> (! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f14, plain, ! [X0] : (! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)) | ~ sP2(X0)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X0] : (sP2(X0) <=> ! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f15, plain, ! [X33] : ((! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59))) | ~ sP3(X33)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X33] : (sP3(X33) <=> (! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f16, plain, ! [X44] : (! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)) | ~ sP4(X44)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X44] : (sP4(X44) <=> ! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f17, plain, ! [X34] : ((! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40))) | ~ sP5(X34)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X34] : (sP5(X34) <=> (! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f18, plain, ! [X34] : (! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | sP4(X44))) | ~ r1(X34, X44)) | ~ sP6(X34)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X34] : (sP6(X34) <=> ! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | sP4(X44))) | ~ r1(X34, X44))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f19, plain, ! [X0] : (((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | sP2(X0))) | ~ sP7(X0)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X0] : (sP7(X0) <=> ((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | sP2(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f20, plain, ! [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ~ sP8(X26)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X26] : (sP8(X26) <=> ! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f21, plain, ! [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ~ sP9(X19)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X19] : (sP9(X19) <=> ! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f11, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : ((! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p2(X34)) | (! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40))) | ! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | ! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) & ((! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) & ~ p2(X33)) | (! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59)))) & r1(X0, X33)) | ((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | ! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68))))) & ! [X72] : (? [X73] : (p1(X73) & ? [X74] : (~ p1(X74) & r1(X73, X74)) & r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) & ? [X75] : (~ p1(X75) & r1(X0, X75)) & ((! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79))) | ! [X80] : (? [X81] : (p5(X81) & r1(X80, X81)) | ~ r1(X0, X80))) & ! [X82] : (? [X83] : (p3(X83) & ? [X84] : (~ p3(X84) & r1(X83, X84)) & r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) & ? [X85] : (~ p3(X85) & r1(X0, X85)) & ((! [X86] : (? [X87] : (p2(X87) & ? [X88] : (~ p2(X88) & r1(X87, X88)) & r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) & ? [X89] : (~ p2(X89) & r1(X0, X89))) | ? [X90] : (? [X91] : (! [X92] : (~ p5(X92) | ~ r1(X91, X92)) & r1(X90, X91)) & ~ p1(X90) & r1(X0, X90)))), inference(flattening, [], [f10])).
fof(f10, plain, ? [X0] : ((! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : (((! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p2(X34)) | (! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40)))) | ! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | ! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) & ((! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) & ~ p2(X33)) | (! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59)))) & r1(X0, X33)) | ((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | ! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) & ! [X72] : (? [X73] : (p1(X73) & ? [X74] : (~ p1(X74) & r1(X73, X74)) & r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) & ? [X75] : (~ p1(X75) & r1(X0, X75)) & ((! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79))) | ! [X80] : (? [X81] : (p5(X81) & r1(X80, X81)) | ~ r1(X0, X80))) & ! [X82] : (? [X83] : (p3(X83) & ? [X84] : (~ p3(X84) & r1(X83, X84)) & r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) & ? [X85] : (~ p3(X85) & r1(X0, X85)) & ((! [X86] : (? [X87] : (p2(X87) & ? [X88] : (~ p2(X88) & r1(X87, X88)) & r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) & ? [X89] : (~ p2(X89) & r1(X0, X89))) | ? [X90] : (? [X91] : (! [X92] : (~ p5(X92) | ~ r1(X91, X92)) & r1(X90, X91)) & ~ p1(X90) & r1(X0, X90)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ~ r1(X6, X7) | p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ~ r1(X5, X10) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ~ r1(X13, X14) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ~ r1(X12, X17) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ~ r1(X20, X21) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ~ r1(X19, X24) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ~ r1(X27, X28) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ~ r1(X26, X31) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p2(X34)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X34, X40)))) | ! [X44] : (((~ ! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p2(X44)) & (~ ! [X47] : (~ ! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | p2(X47) | ~ r1(X44, X47)) | ! [X50] : (! [X51] : (~ ! [X52] : (~ p2(X52) | ! [X53] : (p2(X53) | ~ r1(X52, X53)) | ~ r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) | ((~ ! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) | p2(X33)) & (~ ! [X56] : (~ ! [X57] : (~ p2(X57) | ! [X58] : (p2(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) | ! [X59] : (! [X60] : (~ ! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | p2(X60) | ~ r1(X59, X60)) | ~ r1(X33, X59)))) | ~ r1(X0, X33)) | ((~ ! [X63] : (~ p2(X63) | ! [X64] : (p2(X64) | ~ r1(X63, X64)) | ~ r1(X0, X63)) | p2(X0)) & (~ ! [X65] : (~ ! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) | p2(X65) | ~ r1(X0, X65)) | ! [X68] : (! [X69] : (~ ! [X70] : (~ p2(X70) | ! [X71] : (p2(X71) | ~ r1(X70, X71)) | ~ r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) | ~ ! [X72] : (~ ! [X73] : (~ p1(X73) | ! [X74] : (p1(X74) | ~ r1(X73, X74)) | ~ r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) | ! [X75] : (p1(X75) | ~ r1(X0, X75)) | ((~ ! [X76] : (~ ! [X77] : (~ p2(X77) | ! [X78] : (p2(X78) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) | ! [X79] : (p2(X79) | ~ r1(X0, X79))) & ~ ! [X80] : (~ ! [X81] : (~ p5(X81) | ~ r1(X80, X81)) | ~ r1(X0, X80))) | ~ ! [X82] : (~ ! [X83] : (~ p3(X83) | ! [X84] : (p3(X84) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) | ! [X85] : (p3(X85) | ~ r1(X0, X85)) | ((~ ! [X86] : (~ ! [X87] : (~ p2(X87) | ! [X88] : (p2(X88) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) | ! [X89] : (p2(X89) | ~ r1(X0, X89))) & ! [X90] : (! [X91] : (~ ! [X92] : (~ p5(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | p1(X90) | ~ r1(X0, X90)))), inference(flattening, [], [f6])).
fof(f6, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ~ r1(X6, X7) | p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ~ r1(X5, X10) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ~ r1(X13, X14) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ~ r1(X12, X17) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ~ r1(X20, X21) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ~ r1(X19, X24) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ~ r1(X27, X28) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ~ r1(X26, X31) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p2(X34)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X34, X40)))) | ! [X44] : (((~ ! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p2(X44)) & (~ ! [X47] : (~ ! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | p2(X47) | ~ r1(X44, X47)) | ! [X50] : (! [X51] : (~ ! [X52] : (~ p2(X52) | ! [X53] : (p2(X53) | ~ r1(X52, X53)) | ~ r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) | ((~ ! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) | p2(X33)) & (~ ! [X56] : (~ ! [X57] : (~ p2(X57) | ! [X58] : (p2(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) | ! [X59] : (! [X60] : (~ ! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | p2(X60) | ~ r1(X59, X60)) | ~ r1(X33, X59)))) | ~ r1(X0, X33)) | ((~ ! [X63] : (~ p2(X63) | ! [X64] : (p2(X64) | ~ r1(X63, X64)) | ~ r1(X0, X63)) | p2(X0)) & (~ ! [X65] : (~ ! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) | p2(X65) | ~ r1(X0, X65)) | ! [X68] : (! [X69] : (~ ! [X70] : (~ p2(X70) | ! [X71] : (p2(X71) | ~ r1(X70, X71)) | ~ r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) | ~ ! [X72] : (~ ! [X73] : (~ p1(X73) | ! [X74] : (p1(X74) | ~ r1(X73, X74)) | ~ r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) | ! [X75] : (p1(X75) | ~ r1(X0, X75)) | ((~ ! [X76] : (~ ! [X77] : (~ p2(X77) | ! [X78] : (p2(X78) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) | ! [X79] : (p2(X79) | ~ r1(X0, X79))) & ~ ! [X80] : (~ ! [X81] : (~ p5(X81) | ~ r1(X80, X81)) | ~ r1(X0, X80))) | ~ ! [X82] : (~ ! [X83] : (~ p3(X83) | ! [X84] : (p3(X84) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) | ! [X85] : (p3(X85) | ~ r1(X0, X85)) | ((~ ! [X86] : (~ ! [X87] : (~ p2(X87) | ! [X88] : (p2(X88) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) | ! [X89] : (p2(X89) | ~ r1(X0, X89))) & ! [X90] : (! [X91] : (~ ! [X92] : (~ p5(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | p1(X90) | ~ r1(X0, X90)))), inference(true_and_false_elimination, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ($false | ~ r1(X6, X7)) | p1(X6)) | ! [X8] : (! [X9] : ($false | ~ r1(X8, X9)) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ($false | ~ r1(X5, X10)) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ($false | ~ r1(X0, X11)) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ($false | ~ r1(X13, X14)) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ($false | ~ r1(X15, X16)) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ($false | ~ r1(X12, X17)) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ($false | ~ r1(X0, X18)) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ($false | ~ r1(X20, X21)) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ($false | ~ r1(X22, X23)) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ($false | ~ r1(X19, X24)) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ($false | ~ r1(X0, X25)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ($false | ~ r1(X27, X28)) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ($false | ~ r1(X29, X30)) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ($false | ~ r1(X26, X31)) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ($false | ~ r1(X0, X32)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p2(X34)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X34, X40)))) | ! [X44] : (((~ ! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p2(X44)) & (~ ! [X47] : (~ ! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | p2(X47) | ~ r1(X44, X47)) | ! [X50] : (! [X51] : (~ ! [X52] : (~ p2(X52) | ! [X53] : (p2(X53) | ~ r1(X52, X53)) | ~ r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) | ((~ ! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) | p2(X33)) & (~ ! [X56] : (~ ! [X57] : (~ p2(X57) | ! [X58] : (p2(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) | ! [X59] : (! [X60] : (~ ! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | p2(X60) | ~ r1(X59, X60)) | ~ r1(X33, X59)))) | ~ r1(X0, X33)) | ((~ ! [X63] : (~ p2(X63) | ! [X64] : (p2(X64) | ~ r1(X63, X64)) | ~ r1(X0, X63)) | p2(X0)) & (~ ! [X65] : (~ ! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) | p2(X65) | ~ r1(X0, X65)) | ! [X68] : (! [X69] : (~ ! [X70] : (~ p2(X70) | ! [X71] : (p2(X71) | ~ r1(X70, X71)) | ~ r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) | ~ ! [X72] : (~ ! [X73] : (~ p1(X73) | ! [X74] : (p1(X74) | ~ r1(X73, X74)) | ~ r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) | ! [X75] : (p1(X75) | ~ r1(X0, X75)) | ((~ ! [X76] : (~ ! [X77] : (~ p2(X77) | ! [X78] : (p2(X78) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) | ! [X79] : (p2(X79) | ~ r1(X0, X79))) & ~ ! [X80] : (~ ! [X81] : (~ p5(X81) | ~ r1(X80, X81)) | ~ r1(X0, X80))) | ~ ! [X82] : (~ ! [X83] : (~ p3(X83) | ! [X84] : (p3(X84) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | p3(X82) | ~ r1(X0, X82)) | ! [X85] : (p3(X85) | ~ r1(X0, X85)) | ((~ ! [X86] : (~ ! [X87] : (~ p2(X87) | ! [X88] : (p2(X88) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) | ! [X89] : (p2(X89) | ~ r1(X0, X89))) & ! [X90] : (! [X91] : (~ ! [X92] : (~ p5(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | p1(X90) | ~ r1(X0, X90)))), inference(rectify, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ ! [X0] : (~ p5(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ! [X1] : (! [X0] : (~ ! [X1] : (~ p5(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ ! [X0] : (~ p5(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ! [X1] : (! [X0] : (~ ! [X1] : (~ p5(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL676+1.005.p', main)).
fof(f177765, plain, (p2(sK37(sK32(sK36))) | ~ spl58_4715), inference(avatar_component_clause, [], [f177763])).
fof(f177763, plain, (spl58_4715 <=> p2(sK37(sK32(sK36)))), introduced(avatar_definition, [new_symbols(naming, [spl58_4715])])).
fof(f201167, plain, (spl58_4702 | ~ spl58_171 | spl58_771 | ~ spl58_5402), inference(avatar_split_clause, [], [f201166, f201159, f23148, f2893, f177687])).
fof(f177687, plain, (spl58_4702 <=> ! [X167] : (~ r1(X167, sK32(sK36)) | ~ sP1(X167))), introduced(avatar_definition, [new_symbols(naming, [spl58_4702])])).
fof(f201159, plain, (spl58_5402 <=> ! [X0] : (~ r1(X0, sK37(sK32(sK36))) | ~ sP1(X0))), introduced(avatar_definition, [new_symbols(naming, [spl58_5402])])).
fof(f201166, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK32(sK36))) | (~ spl58_171 | spl58_771 | ~ spl58_5402)), inference(subsumption_resolution, [], [f201165, f2894])).
fof(f201165, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK32(sK36)) | ~ r1(sK36, sK32(sK36))) | (spl58_771 | ~ spl58_5402)), inference(subsumption_resolution, [], [f201163, f23149])).
fof(f201163, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK32(sK36)) | p2(sK32(sK36)) | ~ r1(sK36, sK32(sK36))) | ~ spl58_5402), inference(resolution, [], [f201160, f850])).
fof(f850, plain, ! [X14, X15] : (r1(X14, sK37(X15)) | ~ r1(X14, X15) | p2(X15) | ~ r1(sK36, X15)), inference(resolution, [], [f104, f212])).
fof(f212, plain, ! [X1] : (r1(X1, sK37(X1)) | p2(X1) | ~ r1(sK36, X1)), inference(cnf_transformation, [], [f102])).
fof(f104, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | r1(X0, X2) | ~ r1(X0, X1)), inference(cnf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1, X2] : (r1(X0, X2) | ~ r1(X1, X2) | ~ r1(X0, X1)), inference(flattening, [], [f8])).
fof(f8, plain, ! [X0, X1, X2] : (r1(X0, X2) | (~ r1(X1, X2) | ~ r1(X0, X1))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1, X2] : ((r1(X1, X2) & r1(X0, X1)) => r1(X0, X2)), file('/home/ubuntu/library/tptp/Problems/LCL/LCL676+1.005.p', transitivity)).
fof(f201160, plain, (! [X0] : (~ r1(X0, sK37(sK32(sK36))) | ~ sP1(X0)) | ~ spl58_5402), inference(avatar_component_clause, [], [f201159])).
fof(f201161, plain, (spl58_5402 | spl58_4715 | ~ spl58_171 | spl58_771 | ~ spl58_4701 | ~ spl58_4706 | ~ spl58_4711), inference(avatar_split_clause, [], [f201157, f177737, f177709, f177683, f23148, f2893, f177763, f201159])).
fof(f177683, plain, (spl58_4701 <=> r1(sK30(sK37(sK32(sK36))), sK31(sK37(sK32(sK36))))), introduced(avatar_definition, [new_symbols(naming, [spl58_4701])])).
fof(f177709, plain, (spl58_4706 <=> r1(sK37(sK32(sK36)), sK30(sK37(sK32(sK36))))), introduced(avatar_definition, [new_symbols(naming, [spl58_4706])])).
fof(f177737, plain, (spl58_4711 <=> p2(sK30(sK37(sK32(sK36))))), introduced(avatar_definition, [new_symbols(naming, [spl58_4711])])).
fof(f201157, plain, (! [X0] : (p2(sK37(sK32(sK36))) | ~ r1(X0, sK37(sK32(sK36))) | ~ sP1(X0)) | (~ spl58_171 | spl58_771 | ~ spl58_4701 | ~ spl58_4706 | ~ spl58_4711)), inference(resolution, [], [f200748, f156])).
fof(f156, plain, ! [X0, X1] : (~ p2(sK31(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f72])).
fof(f72, plain, ! [X0] : ((! [X1] : ((p2(sK30(X1)) & (~ p2(sK31(X1)) & r1(sK30(X1), sK31(X1))) & r1(X1, sK30(X1))) | p2(X1) | ~ r1(X0, X1)) & (~ p2(sK32(X0)) & r1(X0, sK32(X0)))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK30, sK31, sK32])], [f68, f71, f70, f69])).
fof(f69, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK30(X1)) & ? [X3] : (~ p2(X3) & r1(sK30(X1), X3)) & r1(X1, sK30(X1)))), introduced(choice_axiom, [])).
fof(f70, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK30(X1), X3)) => (~ p2(sK31(X1)) & r1(sK30(X1), sK31(X1)))), introduced(choice_axiom, [])).
fof(f71, plain, ! [X0] : (? [X4] : (~ p2(X4) & r1(X0, X4)) => (~ p2(sK32(X0)) & r1(X0, sK32(X0)))), introduced(choice_axiom, [])).
fof(f68, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (~ p2(X4) & r1(X0, X4))) | ~ sP1(X0)), inference(rectify, [], [f67])).
fof(f67, plain, ! [X0] : ((! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79))) | ~ sP1(X0)), inference(nnf_transformation, [], [f13])).
fof(f200748, plain, (p2(sK31(sK37(sK32(sK36)))) | (~ spl58_171 | spl58_771 | ~ spl58_4701 | ~ spl58_4706 | ~ spl58_4711)), inference(resolution, [], [f177685, f194898])).
fof(f194898, plain, (! [X0] : (~ r1(sK30(sK37(sK32(sK36))), X0) | p2(X0)) | (~ spl58_171 | spl58_771 | ~ spl58_4706 | ~ spl58_4711)), inference(subsumption_resolution, [], [f194897, f2894])).
fof(f194897, plain, (! [X0] : (p2(X0) | ~ r1(sK30(sK37(sK32(sK36))), X0) | ~ r1(sK36, sK32(sK36))) | (spl58_771 | ~ spl58_4706 | ~ spl58_4711)), inference(subsumption_resolution, [], [f194896, f23149])).
fof(f194896, plain, (! [X0] : (p2(X0) | ~ r1(sK30(sK37(sK32(sK36))), X0) | p2(sK32(sK36)) | ~ r1(sK36, sK32(sK36))) | (~ spl58_4706 | ~ spl58_4711)), inference(subsumption_resolution, [], [f194705, f177739])).
fof(f177739, plain, (p2(sK30(sK37(sK32(sK36)))) | ~ spl58_4711), inference(avatar_component_clause, [], [f177737])).
fof(f194705, plain, (! [X0] : (p2(X0) | ~ r1(sK30(sK37(sK32(sK36))), X0) | ~ p2(sK30(sK37(sK32(sK36)))) | p2(sK32(sK36)) | ~ r1(sK36, sK32(sK36))) | ~ spl58_4706), inference(resolution, [], [f177711, f214])).
fof(f214, plain, ! [X4, X3, X1] : (~ r1(sK37(X1), X3) | p2(X4) | ~ r1(X3, X4) | ~ p2(X3) | p2(X1) | ~ r1(sK36, X1)), inference(cnf_transformation, [], [f102])).
fof(f177711, plain, (r1(sK37(sK32(sK36)), sK30(sK37(sK32(sK36)))) | ~ spl58_4706), inference(avatar_component_clause, [], [f177709])).
fof(f177685, plain, (r1(sK30(sK37(sK32(sK36))), sK31(sK37(sK32(sK36)))) | ~ spl58_4701), inference(avatar_component_clause, [], [f177683])).
fof(f185469, plain, (~ spl58_39 | ~ spl58_4702), inference(avatar_contradiction_clause, [], [f185468])).
fof(f185468, plain, ($false | (~ spl58_39 | ~ spl58_4702)), inference(subsumption_resolution, [], [f185467, f390])).
fof(f390, plain, (sP1(sK36) | ~ spl58_39), inference(avatar_component_clause, [], [f388])).
fof(f388, plain, (spl58_39 <=> sP1(sK36)), introduced(avatar_definition, [new_symbols(naming, [spl58_39])])).
fof(f185467, plain, (~ sP1(sK36) | ~ spl58_4702), inference(duplicate_literal_removal, [], [f185454])).
fof(f185454, plain, (~ sP1(sK36) | ~ sP1(sK36) | ~ spl58_4702), inference(resolution, [], [f177688, f152])).
fof(f152, plain, ! [X0] : (r1(X0, sK32(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f72])).
fof(f177688, plain, (! [X167] : (~ r1(X167, sK32(sK36)) | ~ sP1(X167)) | ~ spl58_4702), inference(avatar_component_clause, [], [f177687])).
fof(f177740, plain, (spl58_4711 | spl58_4702 | spl58_771 | ~ spl58_39), inference(avatar_split_clause, [], [f177735, f388, f23148, f177687, f177737])).
fof(f177735, plain, (! [X77] : (p2(sK32(sK36)) | ~ r1(X77, sK32(sK36)) | p2(sK30(sK37(sK32(sK36)))) | ~ sP1(X77)) | ~ spl58_39), inference(subsumption_resolution, [], [f56459, f390])).
fof(f56459, plain, ! [X77] : (p2(sK32(sK36)) | ~ r1(X77, sK32(sK36)) | p2(sK30(sK37(sK32(sK36)))) | ~ sP1(X77) | ~ sP1(sK36)), inference(resolution, [], [f6263, f152])).
fof(f6263, plain, ! [X81, X82] : (~ r1(sK36, X82) | p2(X82) | ~ r1(X81, X82) | p2(sK30(sK37(X82))) | ~ sP1(X81)), inference(subsumption_resolution, [], [f6217, f213])).
fof(f6217, plain, ! [X81, X82] : (~ r1(X81, X82) | p2(X82) | ~ r1(sK36, X82) | p2(sK37(X82)) | p2(sK30(sK37(X82))) | ~ sP1(X81)), inference(resolution, [], [f850, f157])).
fof(f157, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK30(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f72])).
fof(f177712, plain, (spl58_4706 | spl58_4702 | spl58_771 | ~ spl58_39), inference(avatar_split_clause, [], [f177707, f388, f23148, f177687, f177709])).
fof(f177707, plain, (! [X127] : (p2(sK32(sK36)) | ~ r1(X127, sK32(sK36)) | r1(sK37(sK32(sK36)), sK30(sK37(sK32(sK36)))) | ~ sP1(X127)) | ~ spl58_39), inference(subsumption_resolution, [], [f97772, f390])).
fof(f97772, plain, ! [X127] : (p2(sK32(sK36)) | ~ r1(X127, sK32(sK36)) | r1(sK37(sK32(sK36)), sK30(sK37(sK32(sK36)))) | ~ sP1(X127) | ~ sP1(sK36)), inference(resolution, [], [f6261, f152])).
fof(f6261, plain, ! [X78, X77] : (~ r1(sK36, X78) | p2(X78) | ~ r1(X77, X78) | r1(sK37(X78), sK30(sK37(X78))) | ~ sP1(X77)), inference(subsumption_resolution, [], [f6215, f213])).
fof(f6215, plain, ! [X78, X77] : (~ r1(X77, X78) | p2(X78) | ~ r1(sK36, X78) | p2(sK37(X78)) | r1(sK37(X78), sK30(sK37(X78))) | ~ sP1(X77)), inference(resolution, [], [f850, f154])).
fof(f154, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK30(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f72])).
fof(f177689, plain, (spl58_4701 | spl58_4702 | spl58_771 | ~ spl58_39), inference(avatar_split_clause, [], [f177681, f388, f23148, f177687, f177683])).
fof(f177681, plain, (! [X167] : (p2(sK32(sK36)) | ~ r1(X167, sK32(sK36)) | r1(sK30(sK37(sK32(sK36))), sK31(sK37(sK32(sK36)))) | ~ sP1(X167)) | ~ spl58_39), inference(subsumption_resolution, [], [f141709, f390])).
fof(f141709, plain, ! [X167] : (p2(sK32(sK36)) | ~ r1(X167, sK32(sK36)) | r1(sK30(sK37(sK32(sK36))), sK31(sK37(sK32(sK36)))) | ~ sP1(X167) | ~ sP1(sK36)), inference(resolution, [], [f6262, f152])).
fof(f6262, plain, ! [X80, X79] : (~ r1(sK36, X80) | p2(X80) | ~ r1(X79, X80) | r1(sK30(sK37(X80)), sK31(sK37(X80))) | ~ sP1(X79)), inference(subsumption_resolution, [], [f6216, f213])).
fof(f6216, plain, ! [X80, X79] : (~ r1(X79, X80) | p2(X80) | ~ r1(sK36, X80) | p2(sK37(X80)) | r1(sK30(sK37(X80)), sK31(sK37(X80))) | ~ sP1(X79)), inference(resolution, [], [f850, f155])).
fof(f155, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK30(X1), sK31(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f72])).
fof(f177478, plain, (~ spl58_42 | ~ spl58_1839), inference(avatar_split_clause, [], [f177176, f56621, f400])).
fof(f400, plain, (spl58_42 <=> sP0(sK36)), introduced(avatar_definition, [new_symbols(naming, [spl58_42])])).
fof(f56621, plain, (spl58_1839 <=> ! [X4] : (~ r1(X4, sK35(sK36)) | ~ sP0(X4))), introduced(avatar_definition, [new_symbols(naming, [spl58_1839])])).
fof(f177176, plain, (~ sP0(sK36) | ~ spl58_1839), inference(duplicate_literal_removal, [], [f177163])).
fof(f177163, plain, (~ sP0(sK36) | ~ sP0(sK36) | ~ spl58_1839), inference(resolution, [], [f56622, f158])).
fof(f158, plain, ! [X0] : (r1(X0, sK35(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f78])).
fof(f78, plain, ! [X0] : ((! [X1] : ((p2(sK33(X1)) & (~ p2(sK34(X1)) & r1(sK33(X1), sK34(X1))) & r1(X1, sK33(X1))) | p2(X1) | ~ r1(X0, X1)) & (~ p2(sK35(X0)) & r1(X0, sK35(X0)))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK33, sK34, sK35])], [f74, f77, f76, f75])).
fof(f75, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK33(X1)) & ? [X3] : (~ p2(X3) & r1(sK33(X1), X3)) & r1(X1, sK33(X1)))), introduced(choice_axiom, [])).
fof(f76, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK33(X1), X3)) => (~ p2(sK34(X1)) & r1(sK33(X1), sK34(X1)))), introduced(choice_axiom, [])).
fof(f77, plain, ! [X0] : (? [X4] : (~ p2(X4) & r1(X0, X4)) => (~ p2(sK35(X0)) & r1(X0, sK35(X0)))), introduced(choice_axiom, [])).
fof(f74, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (~ p2(X4) & r1(X0, X4))) | ~ sP0(X0)), inference(rectify, [], [f73])).
fof(f73, plain, ! [X0] : ((! [X86] : (? [X87] : (p2(X87) & ? [X88] : (~ p2(X88) & r1(X87, X88)) & r1(X86, X87)) | p2(X86) | ~ r1(X0, X86)) & ? [X89] : (~ p2(X89) & r1(X0, X89))) | ~ sP0(X0)), inference(nnf_transformation, [], [f12])).
fof(f56622, plain, (! [X4] : (~ r1(X4, sK35(sK36)) | ~ sP0(X4)) | ~ spl58_1839), inference(avatar_component_clause, [], [f56621])).
fof(f177473, plain, (~ spl58_39 | ~ spl58_771), inference(avatar_split_clause, [], [f24043, f23148, f388])).
fof(f24043, plain, (~ sP1(sK36) | ~ spl58_771), inference(resolution, [], [f23150, f153])).
fof(f153, plain, ! [X0] : (~ p2(sK32(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f72])).
fof(f23150, plain, (p2(sK32(sK36)) | ~ spl58_771), inference(avatar_component_clause, [], [f23148])).
fof(f176976, plain, (spl58_1839 | spl58_365 | ~ spl58_1041 | spl58_1584 | ~ spl58_1838 | ~ spl58_3076 | ~ spl58_3507), inference(avatar_split_clause, [], [f176975, f141954, f97979, f56617, f44628, f31367, f11496, f56621])).
fof(f11496, plain, (spl58_365 <=> p2(sK35(sK36))), introduced(avatar_definition, [new_symbols(naming, [spl58_365])])).
fof(f31367, plain, (spl58_1041 <=> r1(sK36, sK35(sK36))), introduced(avatar_definition, [new_symbols(naming, [spl58_1041])])).
fof(f44628, plain, (spl58_1584 <=> p2(sK37(sK35(sK36)))), introduced(avatar_definition, [new_symbols(naming, [spl58_1584])])).
fof(f56617, plain, (spl58_1838 <=> p2(sK33(sK37(sK35(sK36))))), introduced(avatar_definition, [new_symbols(naming, [spl58_1838])])).
fof(f97979, plain, (spl58_3076 <=> r1(sK37(sK35(sK36)), sK33(sK37(sK35(sK36))))), introduced(avatar_definition, [new_symbols(naming, [spl58_3076])])).
fof(f141954, plain, (spl58_3507 <=> r1(sK33(sK37(sK35(sK36))), sK34(sK37(sK35(sK36))))), introduced(avatar_definition, [new_symbols(naming, [spl58_3507])])).
fof(f176975, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK35(sK36))) | (spl58_365 | ~ spl58_1041 | spl58_1584 | ~ spl58_1838 | ~ spl58_3076 | ~ spl58_3507)), inference(subsumption_resolution, [], [f176974, f31368])).
fof(f31368, plain, (r1(sK36, sK35(sK36)) | ~ spl58_1041), inference(avatar_component_clause, [], [f31367])).
fof(f176974, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK35(sK36)) | ~ r1(sK36, sK35(sK36))) | (spl58_365 | ~ spl58_1041 | spl58_1584 | ~ spl58_1838 | ~ spl58_3076 | ~ spl58_3507)), inference(subsumption_resolution, [], [f176972, f11497])).
fof(f11497, plain, (~ p2(sK35(sK36)) | spl58_365), inference(avatar_component_clause, [], [f11496])).
fof(f176972, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK35(sK36)) | p2(sK35(sK36)) | ~ r1(sK36, sK35(sK36))) | (spl58_365 | ~ spl58_1041 | spl58_1584 | ~ spl58_1838 | ~ spl58_3076 | ~ spl58_3507)), inference(resolution, [], [f175864, f850])).
fof(f175864, plain, (! [X0] : (~ r1(X0, sK37(sK35(sK36))) | ~ sP0(X0)) | (spl58_365 | ~ spl58_1041 | spl58_1584 | ~ spl58_1838 | ~ spl58_3076 | ~ spl58_3507)), inference(subsumption_resolution, [], [f175863, f44629])).
fof(f44629, plain, (~ p2(sK37(sK35(sK36))) | spl58_1584), inference(avatar_component_clause, [], [f44628])).
fof(f175863, plain, (! [X0] : (p2(sK37(sK35(sK36))) | ~ r1(X0, sK37(sK35(sK36))) | ~ sP0(X0)) | (spl58_365 | ~ spl58_1041 | ~ spl58_1838 | ~ spl58_3076 | ~ spl58_3507)), inference(resolution, [], [f174556, f162])).
fof(f162, plain, ! [X0, X1] : (~ p2(sK34(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f78])).
fof(f174556, plain, (p2(sK34(sK37(sK35(sK36)))) | (spl58_365 | ~ spl58_1041 | ~ spl58_1838 | ~ spl58_3076 | ~ spl58_3507)), inference(resolution, [], [f141956, f101786])).
fof(f101786, plain, (! [X0] : (~ r1(sK33(sK37(sK35(sK36))), X0) | p2(X0)) | (spl58_365 | ~ spl58_1041 | ~ spl58_1838 | ~ spl58_3076)), inference(subsumption_resolution, [], [f101785, f31368])).
fof(f101785, plain, (! [X0] : (p2(X0) | ~ r1(sK33(sK37(sK35(sK36))), X0) | ~ r1(sK36, sK35(sK36))) | (spl58_365 | ~ spl58_1838 | ~ spl58_3076)), inference(subsumption_resolution, [], [f101784, f11497])).
fof(f101784, plain, (! [X0] : (p2(X0) | ~ r1(sK33(sK37(sK35(sK36))), X0) | p2(sK35(sK36)) | ~ r1(sK36, sK35(sK36))) | (~ spl58_1838 | ~ spl58_3076)), inference(subsumption_resolution, [], [f101643, f56619])).
fof(f56619, plain, (p2(sK33(sK37(sK35(sK36)))) | ~ spl58_1838), inference(avatar_component_clause, [], [f56617])).
fof(f101643, plain, (! [X0] : (p2(X0) | ~ r1(sK33(sK37(sK35(sK36))), X0) | ~ p2(sK33(sK37(sK35(sK36)))) | p2(sK35(sK36)) | ~ r1(sK36, sK35(sK36))) | ~ spl58_3076), inference(resolution, [], [f97981, f214])).
fof(f97981, plain, (r1(sK37(sK35(sK36)), sK33(sK37(sK35(sK36)))) | ~ spl58_3076), inference(avatar_component_clause, [], [f97979])).
fof(f141956, plain, (r1(sK33(sK37(sK35(sK36))), sK34(sK37(sK35(sK36)))) | ~ spl58_3507), inference(avatar_component_clause, [], [f141954])).
fof(f141970, plain, (spl58_3507 | spl58_1839 | ~ spl58_42 | spl58_365), inference(avatar_split_clause, [], [f141969, f11496, f400, f56621, f141954])).
fof(f141969, plain, (! [X195] : (~ r1(X195, sK35(sK36)) | r1(sK33(sK37(sK35(sK36))), sK34(sK37(sK35(sK36)))) | ~ sP0(X195)) | (~ spl58_42 | spl58_365)), inference(subsumption_resolution, [], [f141968, f402])).
fof(f402, plain, (sP0(sK36) | ~ spl58_42), inference(avatar_component_clause, [], [f400])).
fof(f141968, plain, (! [X195] : (~ r1(X195, sK35(sK36)) | r1(sK33(sK37(sK35(sK36))), sK34(sK37(sK35(sK36)))) | ~ sP0(X195) | ~ sP0(sK36)) | spl58_365), inference(subsumption_resolution, [], [f141913, f11497])).
fof(f141913, plain, ! [X195] : (p2(sK35(sK36)) | ~ r1(X195, sK35(sK36)) | r1(sK33(sK37(sK35(sK36))), sK34(sK37(sK35(sK36)))) | ~ sP0(X195) | ~ sP0(sK36)), inference(resolution, [], [f6265, f158])).
fof(f6265, plain, ! [X85, X86] : (~ r1(sK36, X86) | p2(X86) | ~ r1(X85, X86) | r1(sK33(sK37(X86)), sK34(sK37(X86))) | ~ sP0(X85)), inference(subsumption_resolution, [], [f6219, f213])).
fof(f6219, plain, ! [X85, X86] : (~ r1(X85, X86) | p2(X86) | ~ r1(sK36, X86) | p2(sK37(X86)) | r1(sK33(sK37(X86)), sK34(sK37(X86))) | ~ sP0(X85)), inference(resolution, [], [f850, f161])).
fof(f161, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK33(X1), sK34(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f78])).
fof(f97999, plain, (spl58_3076 | spl58_1839 | spl58_365 | ~ spl58_1041), inference(avatar_split_clause, [], [f97998, f31367, f11496, f56621, f97979])).
fof(f97998, plain, (! [X150] : (~ r1(X150, sK35(sK36)) | r1(sK37(sK35(sK36)), sK33(sK37(sK35(sK36)))) | ~ sP0(X150)) | (spl58_365 | ~ spl58_1041)), inference(subsumption_resolution, [], [f97997, f103])).
fof(f103, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL676+1.005.p', reflexivity)).
fof(f97997, plain, (! [X150] : (~ r1(X150, sK35(sK36)) | r1(sK37(sK35(sK36)), sK33(sK37(sK35(sK36)))) | ~ sP0(X150) | ~ r1(sK36, sK36)) | (spl58_365 | ~ spl58_1041)), inference(subsumption_resolution, [], [f97944, f11497])).
fof(f97944, plain, (! [X150] : (p2(sK35(sK36)) | ~ r1(X150, sK35(sK36)) | r1(sK37(sK35(sK36)), sK33(sK37(sK35(sK36)))) | ~ sP0(X150) | ~ r1(sK36, sK36)) | ~ spl58_1041), inference(resolution, [], [f6264, f31664])).
fof(f31664, plain, (! [X24] : (r1(X24, sK35(sK36)) | ~ r1(X24, sK36)) | ~ spl58_1041), inference(resolution, [], [f31368, f104])).
fof(f6264, plain, ! [X83, X84] : (~ r1(sK36, X84) | p2(X84) | ~ r1(X83, X84) | r1(sK37(X84), sK33(sK37(X84))) | ~ sP0(X83)), inference(subsumption_resolution, [], [f6218, f213])).
fof(f6218, plain, ! [X83, X84] : (~ r1(X83, X84) | p2(X84) | ~ r1(sK36, X84) | p2(sK37(X84)) | r1(sK37(X84), sK33(sK37(X84))) | ~ sP0(X83)), inference(resolution, [], [f850, f160])).
fof(f160, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK33(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f78])).
fof(f56636, plain, (spl58_1838 | spl58_1839 | spl58_365 | ~ spl58_1041), inference(avatar_split_clause, [], [f56635, f31367, f11496, f56621, f56617])).
fof(f56635, plain, (! [X90] : (~ r1(X90, sK35(sK36)) | p2(sK33(sK37(sK35(sK36)))) | ~ sP0(X90)) | (spl58_365 | ~ spl58_1041)), inference(subsumption_resolution, [], [f56634, f103])).
fof(f56634, plain, (! [X90] : (~ r1(X90, sK35(sK36)) | p2(sK33(sK37(sK35(sK36)))) | ~ sP0(X90) | ~ r1(sK36, sK36)) | (spl58_365 | ~ spl58_1041)), inference(subsumption_resolution, [], [f56585, f11497])).
fof(f56585, plain, (! [X90] : (p2(sK35(sK36)) | ~ r1(X90, sK35(sK36)) | p2(sK33(sK37(sK35(sK36)))) | ~ sP0(X90) | ~ r1(sK36, sK36)) | ~ spl58_1041), inference(resolution, [], [f6266, f31664])).
fof(f6266, plain, ! [X88, X87] : (~ r1(sK36, X88) | p2(X88) | ~ r1(X87, X88) | p2(sK33(sK37(X88))) | ~ sP0(X87)), inference(subsumption_resolution, [], [f6220, f213])).
fof(f6220, plain, ! [X88, X87] : (~ r1(X87, X88) | p2(X88) | ~ r1(sK36, X88) | p2(sK37(X88)) | p2(sK33(sK37(X88))) | ~ sP0(X87)), inference(resolution, [], [f850, f163])).
fof(f163, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK33(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f78])).
fof(f44635, plain, (spl58_365 | ~ spl58_1041 | ~ spl58_1584), inference(avatar_contradiction_clause, [], [f44634])).
fof(f44634, plain, ($false | (spl58_365 | ~ spl58_1041 | ~ spl58_1584)), inference(subsumption_resolution, [], [f44633, f31368])).
fof(f44633, plain, (~ r1(sK36, sK35(sK36)) | (spl58_365 | ~ spl58_1584)), inference(subsumption_resolution, [], [f44632, f11497])).
fof(f44632, plain, (p2(sK35(sK36)) | ~ r1(sK36, sK35(sK36)) | ~ spl58_1584), inference(resolution, [], [f44630, f213])).
fof(f44630, plain, (p2(sK37(sK35(sK36))) | ~ spl58_1584), inference(avatar_component_clause, [], [f44628])).
fof(f31557, plain, (~ spl58_42 | spl58_1041), inference(avatar_contradiction_clause, [], [f31556])).
fof(f31556, plain, ($false | (~ spl58_42 | spl58_1041)), inference(subsumption_resolution, [], [f31555, f402])).
fof(f31555, plain, (~ sP0(sK36) | spl58_1041), inference(subsumption_resolution, [], [f31552, f103])).
fof(f31552, plain, (~ r1(sK36, sK36) | ~ sP0(sK36) | spl58_1041), inference(resolution, [], [f31369, f849])).
fof(f849, plain, ! [X12, X13] : (r1(X12, sK35(X13)) | ~ r1(X12, X13) | ~ sP0(X13)), inference(resolution, [], [f104, f158])).
fof(f31369, plain, (~ r1(sK36, sK35(sK36)) | spl58_1041), inference(avatar_component_clause, [], [f31367])).
fof(f11685, plain, (~ spl58_42 | ~ spl58_365), inference(avatar_contradiction_clause, [], [f11684])).
fof(f11684, plain, ($false | (~ spl58_42 | ~ spl58_365)), inference(subsumption_resolution, [], [f11683, f402])).
fof(f11683, plain, (~ sP0(sK36) | ~ spl58_365), inference(resolution, [], [f11498, f159])).
fof(f159, plain, ! [X0] : (~ p2(sK35(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f78])).
fof(f11498, plain, (p2(sK35(sK36)) | ~ spl58_365), inference(avatar_component_clause, [], [f11496])).
fof(f3033, plain, (~ spl58_101 | ~ spl58_38 | ~ spl58_40 | ~ spl58_41), inference(avatar_split_clause, [], [f3032, f397, f393, f385, f758])).
fof(f758, plain, (spl58_101 <=> r1(sK36, sK57)), introduced(avatar_definition, [new_symbols(naming, [spl58_101])])).
fof(f385, plain, (spl58_38 <=> ! [X35] : (p5(sK52(X35)) | ~ r1(sK36, X35))), introduced(avatar_definition, [new_symbols(naming, [spl58_38])])).
fof(f393, plain, (spl58_40 <=> ! [X35] : (r1(X35, sK52(X35)) | ~ r1(sK36, X35))), introduced(avatar_definition, [new_symbols(naming, [spl58_40])])).
fof(f397, plain, (spl58_41 <=> ! [X43] : (~ p5(X43) | ~ r1(sK57, X43))), introduced(avatar_definition, [new_symbols(naming, [spl58_41])])).
fof(f3032, plain, (~ r1(sK36, sK57) | (~ spl58_38 | ~ spl58_40 | ~ spl58_41)), inference(subsumption_resolution, [], [f1066, f386])).
fof(f386, plain, (! [X35] : (~ r1(sK36, X35) | p5(sK52(X35))) | ~ spl58_38), inference(avatar_component_clause, [], [f385])).
fof(f1066, plain, (~ r1(sK36, sK57) | ~ p5(sK52(sK57)) | (~ spl58_40 | ~ spl58_41)), inference(resolution, [], [f394, f398])).
fof(f398, plain, (! [X43] : (~ r1(sK57, X43) | ~ p5(X43)) | ~ spl58_41), inference(avatar_component_clause, [], [f397])).
fof(f394, plain, (! [X35] : (r1(X35, sK52(X35)) | ~ r1(sK36, X35)) | ~ spl58_40), inference(avatar_component_clause, [], [f393])).
fof(f3031, plain, (spl58_101 | ~ spl58_43 | ~ spl58_45), inference(avatar_split_clause, [], [f3028, f415, f405, f758])).
fof(f405, plain, (spl58_43 <=> r1(sK56, sK57)), introduced(avatar_definition, [new_symbols(naming, [spl58_43])])).
fof(f415, plain, (spl58_45 <=> r1(sK36, sK56)), introduced(avatar_definition, [new_symbols(naming, [spl58_45])])).
fof(f3028, plain, (r1(sK36, sK57) | (~ spl58_43 | ~ spl58_45)), inference(resolution, [], [f2915, f417])).
fof(f417, plain, (r1(sK36, sK56) | ~ spl58_45), inference(avatar_component_clause, [], [f415])).
fof(f2915, plain, (! [X0] : (~ r1(X0, sK56) | r1(X0, sK57)) | ~ spl58_43), inference(resolution, [], [f407, f104])).
fof(f407, plain, (r1(sK56, sK57) | ~ spl58_43), inference(avatar_component_clause, [], [f405])).
fof(f2914, plain, (~ spl58_39 | spl58_171), inference(avatar_split_clause, [], [f2906, f2893, f388])).
fof(f2906, plain, (~ sP1(sK36) | spl58_171), inference(resolution, [], [f2895, f152])).
fof(f2895, plain, (~ r1(sK36, sK32(sK36)) | spl58_171), inference(avatar_component_clause, [], [f2893])).
fof(f418, plain, (spl58_45 | spl58_42), inference(avatar_split_clause, [], [f164, f400, f415])).
fof(f164, plain, (sP0(sK36) | r1(sK36, sK56)), inference(cnf_transformation, [], [f102])).
fof(f408, plain, (spl58_43 | spl58_42), inference(avatar_split_clause, [], [f166, f400, f405])).
fof(f166, plain, (sP0(sK36) | r1(sK56, sK57)), inference(cnf_transformation, [], [f102])).
fof(f403, plain, (spl58_41 | spl58_42), inference(avatar_split_clause, [], [f167, f400, f397])).
fof(f167, plain, ! [X43] : (sP0(sK36) | ~ p5(X43) | ~ r1(sK57, X43)), inference(cnf_transformation, [], [f102])).
fof(f395, plain, (spl58_40 | spl58_39), inference(avatar_split_clause, [], [f174, f388, f393])).
fof(f174, plain, ! [X35] : (sP1(sK36) | r1(X35, sK52(X35)) | ~ r1(sK36, X35)), inference(cnf_transformation, [], [f102])).
fof(f391, plain, (spl58_38 | spl58_39), inference(avatar_split_clause, [], [f175, f388, f385])).
fof(f175, plain, ! [X35] : (sP1(sK36) | p5(sK52(X35)) | ~ r1(sK36, X35)), inference(cnf_transformation, [], [f102])).