fof(f3317, plain, $false, inference(avatar_sat_refutation, [], [f583, f590, f641, f1079, f1130, f1167, f3316])).
fof(f3316, plain, (~ spl23_43 | ~ spl23_117 | ~ spl23_120 | spl23_121), inference(avatar_contradiction_clause, [], [f3315])).
fof(f3315, plain, ($false | (~ spl23_43 | ~ spl23_117 | ~ spl23_120 | spl23_121)), inference(subsumption_resolution, [], [f3314, f998])).
fof(f998, plain, (sP2(sK15(sK22)) | ~ spl23_120), inference(avatar_component_clause, [], [f997])).
fof(f997, plain, (spl23_120 <=> sP2(sK15(sK22))), introduced(avatar_definition, [new_symbols(naming, [spl23_120])])).
fof(f3314, plain, (~ sP2(sK15(sK22)) | (~ spl23_43 | ~ spl23_117 | ~ spl23_120 | spl23_121)), inference(subsumption_resolution, [], [f3313, f987])).
fof(f987, plain, (p101(sK15(sK22)) | ~ spl23_117), inference(avatar_component_clause, [], [f985])).
fof(f985, plain, (spl23_117 <=> p101(sK15(sK22))), introduced(avatar_definition, [new_symbols(naming, [spl23_117])])).
fof(f3313, plain, (~ p101(sK15(sK22)) | ~ sP2(sK15(sK22)) | (~ spl23_43 | ~ spl23_117 | ~ spl23_120 | spl23_121)), inference(subsumption_resolution, [], [f3312, f1002])).
fof(f1002, plain, (~ p102(sK15(sK22)) | spl23_121), inference(avatar_component_clause, [], [f1001])).
fof(f1001, plain, (spl23_121 <=> p102(sK15(sK22))), introduced(avatar_definition, [new_symbols(naming, [spl23_121])])).
fof(f3312, plain, (p102(sK15(sK22)) | ~ p101(sK15(sK22)) | ~ sP2(sK15(sK22)) | (~ spl23_43 | ~ spl23_117 | ~ spl23_120 | spl23_121)), inference(resolution, [], [f2679, f108])).
fof(f108, plain, ! [X0] : (~ p3(sK17(X0)) | p102(X0) | ~ p101(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f50])).
fof(f50, plain, ! [X0] : (~ p101(X0) | p102(X0) | ((p102(sK16(X0)) & ~ p103(sK16(X0)) & p3(sK16(X0)) & r1(X0, sK16(X0))) & (p102(sK17(X0)) & ~ p103(sK17(X0)) & ~ p3(sK17(X0)) & r1(X0, sK17(X0)))) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK16, sK17])], [f47, f49, f48])).
fof(f48, plain, ! [X0] : (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) => (p102(sK16(X0)) & ~ p103(sK16(X0)) & p3(sK16(X0)) & r1(X0, sK16(X0)))), introduced(choice_axiom, [])).
fof(f49, plain, ! [X0] : (? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2)) => (p102(sK17(X0)) & ~ p103(sK17(X0)) & ~ p3(sK17(X0)) & r1(X0, sK17(X0)))), introduced(choice_axiom, [])).
fof(f47, plain, ! [X0] : (~ p101(X0) | p102(X0) | (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) & ? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2))) | ~ sP2(X0)), inference(rectify, [], [f46])).
fof(f46, plain, ! [X5] : (~ p101(X5) | p102(X5) | (? [X20] : (p102(X20) & ~ p103(X20) & p3(X20) & r1(X5, X20)) & ? [X21] : (p102(X21) & ~ p103(X21) & ~ p3(X21) & r1(X5, X21))) | ~ sP2(X5)), inference(nnf_transformation, [], [f11])).
fof(f11, plain, ! [X5] : (~ p101(X5) | p102(X5) | (? [X20] : (p102(X20) & ~ p103(X20) & p3(X20) & r1(X5, X20)) & ? [X21] : (p102(X21) & ~ p103(X21) & ~ p3(X21) & r1(X5, X21))) | ~ sP2(X5)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X5] : (sP2(X5) <=> (~ p101(X5) | p102(X5) | (? [X20] : (p102(X20) & ~ p103(X20) & p3(X20) & r1(X5, X20)) & ? [X21] : (p102(X21) & ~ p103(X21) & ~ p3(X21) & r1(X5, X21))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f2679, plain, (p3(sK17(sK15(sK22))) | (~ spl23_43 | ~ spl23_117 | ~ spl23_120 | spl23_121)), inference(subsumption_resolution, [], [f2678, f998])).
fof(f2678, plain, (p3(sK17(sK15(sK22))) | ~ sP2(sK15(sK22)) | (~ spl23_43 | ~ spl23_117 | spl23_121)), inference(subsumption_resolution, [], [f2677, f987])).
fof(f2677, plain, (p3(sK17(sK15(sK22))) | ~ p101(sK15(sK22)) | ~ sP2(sK15(sK22)) | (~ spl23_43 | spl23_121)), inference(subsumption_resolution, [], [f2672, f1002])).
fof(f2672, plain, (p3(sK17(sK15(sK22))) | p102(sK15(sK22)) | ~ p101(sK15(sK22)) | ~ sP2(sK15(sK22)) | ~ spl23_43), inference(resolution, [], [f640, f107])).
fof(f107, plain, ! [X0] : (r1(X0, sK17(X0)) | p102(X0) | ~ p101(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f50])).
fof(f640, plain, (! [X4] : (~ r1(sK15(sK22), X4) | p3(X4)) | ~ spl23_43), inference(avatar_component_clause, [], [f639])).
fof(f639, plain, (spl23_43 <=> ! [X4] : (~ r1(sK15(sK22), X4) | p3(X4))), introduced(avatar_definition, [new_symbols(naming, [spl23_43])])).
fof(f1167, plain, (~ spl23_6 | ~ spl23_121), inference(avatar_contradiction_clause, [], [f1166])).
fof(f1166, plain, ($false | (~ spl23_6 | ~ spl23_121)), inference(subsumption_resolution, [], [f1165, f334])).
fof(f334, plain, (sP3(sK22) | ~ spl23_6), inference(avatar_component_clause, [], [f333])).
fof(f333, plain, (spl23_6 <=> sP3(sK22)), introduced(avatar_definition, [new_symbols(naming, [spl23_6])])).
fof(f1165, plain, (~ sP3(sK22) | ~ spl23_121), inference(subsumption_resolution, [], [f1164, f134])).
fof(f134, plain, p100(sK22), inference(cnf_transformation, [], [f63])).
fof(f63, plain, (p100(sK22) & ~ p101(sK22) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (sP11(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(sK22, X1)) & ! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (p3(X10) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(sK22, X6))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK22])], [f61, f62])).
fof(f62, plain, (? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (sP11(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (p3(X10) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X0, X6))) => (p100(sK22) & ~ p101(sK22) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (sP11(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(sK22, X1)) & ! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (p3(X10) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(sK22, X6)))), introduced(choice_axiom, [])).
fof(f61, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (sP11(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (p3(X10) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X0, X6))), inference(rectify, [], [f21])).
fof(f21, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (sP11(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (p3(X32) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X0, X28))), inference(definition_folding, [], [f8, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9])).
fof(f9, plain, ! [X5] : (~ p103(X5) | p104(X5) | (? [X24] : (p104(X24) & ~ p105(X24) & p5(X24) & r1(X5, X24)) & ? [X25] : (p104(X25) & ~ p105(X25) & ~ p5(X25) & r1(X5, X25))) | ~ sP0(X5)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X5] : (sP0(X5) <=> (~ p103(X5) | p104(X5) | (? [X24] : (p104(X24) & ~ p105(X24) & p5(X24) & r1(X5, X24)) & ? [X25] : (p104(X25) & ~ p105(X25) & ~ p5(X25) & r1(X5, X25))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f10, plain, ! [X5] : (~ p102(X5) | p103(X5) | (? [X22] : (p103(X22) & ~ p104(X22) & p4(X22) & r1(X5, X22)) & ? [X23] : (p103(X23) & ~ p104(X23) & ~ p4(X23) & r1(X5, X23))) | ~ sP1(X5)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X5] : (sP1(X5) <=> (~ p102(X5) | p103(X5) | (? [X22] : (p103(X22) & ~ p104(X22) & p4(X22) & r1(X5, X22)) & ? [X23] : (p103(X23) & ~ p104(X23) & ~ p4(X23) & r1(X5, X23))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f12, plain, ! [X5] : (~ p100(X5) | p101(X5) | (? [X18] : (p101(X18) & ~ p102(X18) & p2(X18) & r1(X5, X18)) & ? [X19] : (p101(X19) & ~ p102(X19) & ~ p2(X19) & r1(X5, X19))) | ~ sP3(X5)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X5] : (sP3(X5) <=> (~ p100(X5) | p101(X5) | (? [X18] : (p101(X18) & ~ p102(X18) & p2(X18) & r1(X5, X18)) & ? [X19] : (p101(X19) & ~ p102(X19) & ~ p2(X19) & r1(X5, X19))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f13, plain, ! [X5] : (~ p104(X5) | p105(X5) | (? [X26] : (p105(X26) & p6(X26) & r1(X5, X26)) & ? [X27] : (p105(X27) & ~ p6(X27) & r1(X5, X27))) | ~ sP4(X5)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X5] : (sP4(X5) <=> (~ p104(X5) | p105(X5) | (? [X26] : (p105(X26) & p6(X26) & r1(X5, X26)) & ? [X27] : (p105(X27) & ~ p6(X27) & r1(X5, X27))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f14, plain, ! [X5] : (~ p105(X5) | ((~ p6(X5) | ! [X16] : (~ p105(X16) | p6(X16) | ~ r1(X5, X16))) & (p6(X5) | ! [X17] : (~ p105(X17) | ~ p6(X17) | ~ r1(X5, X17)))) | ~ sP5(X5)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X5] : (sP5(X5) <=> (~ p105(X5) | ((~ p6(X5) | ! [X16] : (~ p105(X16) | p6(X16) | ~ r1(X5, X16))) & (p6(X5) | ! [X17] : (~ p105(X17) | ~ p6(X17) | ~ r1(X5, X17)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f15, plain, ! [X5] : (~ p104(X5) | ((~ p5(X5) | ! [X14] : (~ p104(X14) | p5(X14) | ~ r1(X5, X14))) & (p5(X5) | ! [X15] : (~ p104(X15) | ~ p5(X15) | ~ r1(X5, X15)))) | ~ sP6(X5)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X5] : (sP6(X5) <=> (~ p104(X5) | ((~ p5(X5) | ! [X14] : (~ p104(X14) | p5(X14) | ~ r1(X5, X14))) & (p5(X5) | ! [X15] : (~ p104(X15) | ~ p5(X15) | ~ r1(X5, X15)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f16, plain, ! [X5] : (~ p103(X5) | ((~ p4(X5) | ! [X12] : (~ p103(X12) | p4(X12) | ~ r1(X5, X12))) & (p4(X5) | ! [X13] : (~ p103(X13) | ~ p4(X13) | ~ r1(X5, X13)))) | ~ sP7(X5)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X5] : (sP7(X5) <=> (~ p103(X5) | ((~ p4(X5) | ! [X12] : (~ p103(X12) | p4(X12) | ~ r1(X5, X12))) & (p4(X5) | ! [X13] : (~ p103(X13) | ~ p4(X13) | ~ r1(X5, X13)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f17, plain, ! [X5] : (~ p102(X5) | ((~ p3(X5) | ! [X10] : (~ p102(X10) | p3(X10) | ~ r1(X5, X10))) & (p3(X5) | ! [X11] : (~ p102(X11) | ~ p3(X11) | ~ r1(X5, X11)))) | ~ sP8(X5)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X5] : (sP8(X5) <=> (~ p102(X5) | ((~ p3(X5) | ! [X10] : (~ p102(X10) | p3(X10) | ~ r1(X5, X10))) & (p3(X5) | ! [X11] : (~ p102(X11) | ~ p3(X11) | ~ r1(X5, X11)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f18, plain, ! [X5] : (~ p101(X5) | ((~ p2(X5) | ! [X8] : (~ p101(X8) | p2(X8) | ~ r1(X5, X8))) & (p2(X5) | ! [X9] : (~ p101(X9) | ~ p2(X9) | ~ r1(X5, X9)))) | ~ sP9(X5)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X5] : (sP9(X5) <=> (~ p101(X5) | ((~ p2(X5) | ! [X8] : (~ p101(X8) | p2(X8) | ~ r1(X5, X8))) & (p2(X5) | ! [X9] : (~ p101(X9) | ~ p2(X9) | ~ r1(X5, X9)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f19, plain, ! [X5] : (~ p100(X5) | ((~ p1(X5) | ! [X6] : (~ p100(X6) | p1(X6) | ~ r1(X5, X6))) & (p1(X5) | ! [X7] : (~ p100(X7) | ~ p1(X7) | ~ r1(X5, X7)))) | ~ sP10(X5)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X5] : (sP10(X5) <=> (~ p100(X5) | ((~ p1(X5) | ! [X6] : (~ p100(X6) | p1(X6) | ~ r1(X5, X6))) & (p1(X5) | ! [X7] : (~ p100(X7) | ~ p1(X7) | ~ r1(X5, X7)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f20, plain, ! [X5] : (((~ p101(X5) | p100(X5)) & (~ p102(X5) | p101(X5)) & (~ p103(X5) | p102(X5)) & (~ p104(X5) | p103(X5)) & (~ p105(X5) | p104(X5)) & sP10(X5) & sP9(X5) & sP8(X5) & sP7(X5) & sP6(X5) & sP5(X5) & sP3(X5) & sP2(X5) & sP1(X5) & sP0(X5) & sP4(X5)) | ~ sP11(X5)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X5] : (sP11(X5) <=> ((~ p101(X5) | p100(X5)) & (~ p102(X5) | p101(X5)) & (~ p103(X5) | p102(X5)) & (~ p104(X5) | p103(X5)) & (~ p105(X5) | p104(X5)) & sP10(X5) & sP9(X5) & sP8(X5) & sP7(X5) & sP6(X5) & sP5(X5) & sP3(X5) & sP2(X5) & sP1(X5) & sP0(X5) & sP4(X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f8, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (((~ p101(X5) | p100(X5)) & (~ p102(X5) | p101(X5)) & (~ p103(X5) | p102(X5)) & (~ p104(X5) | p103(X5)) & (~ p105(X5) | p104(X5)) & (~ p100(X5) | ((~ p1(X5) | ! [X6] : (~ p100(X6) | p1(X6) | ~ r1(X5, X6))) & (p1(X5) | ! [X7] : (~ p100(X7) | ~ p1(X7) | ~ r1(X5, X7))))) & (~ p101(X5) | ((~ p2(X5) | ! [X8] : (~ p101(X8) | p2(X8) | ~ r1(X5, X8))) & (p2(X5) | ! [X9] : (~ p101(X9) | ~ p2(X9) | ~ r1(X5, X9))))) & (~ p102(X5) | ((~ p3(X5) | ! [X10] : (~ p102(X10) | p3(X10) | ~ r1(X5, X10))) & (p3(X5) | ! [X11] : (~ p102(X11) | ~ p3(X11) | ~ r1(X5, X11))))) & (~ p103(X5) | ((~ p4(X5) | ! [X12] : (~ p103(X12) | p4(X12) | ~ r1(X5, X12))) & (p4(X5) | ! [X13] : (~ p103(X13) | ~ p4(X13) | ~ r1(X5, X13))))) & (~ p104(X5) | ((~ p5(X5) | ! [X14] : (~ p104(X14) | p5(X14) | ~ r1(X5, X14))) & (p5(X5) | ! [X15] : (~ p104(X15) | ~ p5(X15) | ~ r1(X5, X15))))) & (~ p105(X5) | ((~ p6(X5) | ! [X16] : (~ p105(X16) | p6(X16) | ~ r1(X5, X16))) & (p6(X5) | ! [X17] : (~ p105(X17) | ~ p6(X17) | ~ r1(X5, X17))))) & (~ p100(X5) | p101(X5) | (? [X18] : (p101(X18) & ~ p102(X18) & p2(X18) & r1(X5, X18)) & ? [X19] : (p101(X19) & ~ p102(X19) & ~ p2(X19) & r1(X5, X19)))) & (~ p101(X5) | p102(X5) | (? [X20] : (p102(X20) & ~ p103(X20) & p3(X20) & r1(X5, X20)) & ? [X21] : (p102(X21) & ~ p103(X21) & ~ p3(X21) & r1(X5, X21)))) & (~ p102(X5) | p103(X5) | (? [X22] : (p103(X22) & ~ p104(X22) & p4(X22) & r1(X5, X22)) & ? [X23] : (p103(X23) & ~ p104(X23) & ~ p4(X23) & r1(X5, X23)))) & (~ p103(X5) | p104(X5) | (? [X24] : (p104(X24) & ~ p105(X24) & p5(X24) & r1(X5, X24)) & ? [X25] : (p104(X25) & ~ p105(X25) & ~ p5(X25) & r1(X5, X25)))) & (~ p104(X5) | p105(X5) | (? [X26] : (p105(X26) & p6(X26) & r1(X5, X26)) & ? [X27] : (p105(X27) & ~ p6(X27) & r1(X5, X27))))) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (p3(X32) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X0, X28))), inference(flattening, [], [f7])).
fof(f7, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (((~ p101(X5) | p100(X5)) & (~ p102(X5) | p101(X5)) & (~ p103(X5) | p102(X5)) & (~ p104(X5) | p103(X5)) & (~ p105(X5) | p104(X5)) & (~ p100(X5) | ((~ p1(X5) | ! [X6] : (~ p100(X6) | p1(X6) | ~ r1(X5, X6))) & (p1(X5) | ! [X7] : (~ p100(X7) | ~ p1(X7) | ~ r1(X5, X7))))) & (~ p101(X5) | ((~ p2(X5) | ! [X8] : (~ p101(X8) | p2(X8) | ~ r1(X5, X8))) & (p2(X5) | ! [X9] : (~ p101(X9) | ~ p2(X9) | ~ r1(X5, X9))))) & (~ p102(X5) | ((~ p3(X5) | ! [X10] : (~ p102(X10) | p3(X10) | ~ r1(X5, X10))) & (p3(X5) | ! [X11] : (~ p102(X11) | ~ p3(X11) | ~ r1(X5, X11))))) & (~ p103(X5) | ((~ p4(X5) | ! [X12] : (~ p103(X12) | p4(X12) | ~ r1(X5, X12))) & (p4(X5) | ! [X13] : (~ p103(X13) | ~ p4(X13) | ~ r1(X5, X13))))) & (~ p104(X5) | ((~ p5(X5) | ! [X14] : (~ p104(X14) | p5(X14) | ~ r1(X5, X14))) & (p5(X5) | ! [X15] : (~ p104(X15) | ~ p5(X15) | ~ r1(X5, X15))))) & (~ p105(X5) | ((~ p6(X5) | ! [X16] : (~ p105(X16) | p6(X16) | ~ r1(X5, X16))) & (p6(X5) | ! [X17] : (~ p105(X17) | ~ p6(X17) | ~ r1(X5, X17))))) & ((~ p100(X5) | p101(X5)) | (? [X18] : ((p101(X18) & ~ p102(X18) & p2(X18)) & r1(X5, X18)) & ? [X19] : ((p101(X19) & ~ p102(X19) & ~ p2(X19)) & r1(X5, X19)))) & ((~ p101(X5) | p102(X5)) | (? [X20] : ((p102(X20) & ~ p103(X20) & p3(X20)) & r1(X5, X20)) & ? [X21] : ((p102(X21) & ~ p103(X21) & ~ p3(X21)) & r1(X5, X21)))) & ((~ p102(X5) | p103(X5)) | (? [X22] : ((p103(X22) & ~ p104(X22) & p4(X22)) & r1(X5, X22)) & ? [X23] : ((p103(X23) & ~ p104(X23) & ~ p4(X23)) & r1(X5, X23)))) & ((~ p103(X5) | p104(X5)) | (? [X24] : ((p104(X24) & ~ p105(X24) & p5(X24)) & r1(X5, X24)) & ? [X25] : ((p104(X25) & ~ p105(X25) & ~ p5(X25)) & r1(X5, X25)))) & ((~ p104(X5) | p105(X5)) | (? [X26] : ((p105(X26) & p6(X26)) & r1(X5, X26)) & ? [X27] : ((p105(X27) & ~ p6(X27)) & r1(X5, X27))))) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) & ! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (p3(X32) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X0, X28))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (((~ p101(X5) | p100(X5)) & (~ p102(X5) | p101(X5)) & (~ p103(X5) | p102(X5)) & (~ p104(X5) | p103(X5)) & (~ p105(X5) | p104(X5)) & (~ p100(X5) | ((~ p1(X5) | ! [X6] : (~ p100(X6) | p1(X6) | ~ r1(X5, X6))) & (p1(X5) | ! [X7] : (~ p100(X7) | ~ p1(X7) | ~ r1(X5, X7))))) & (~ p101(X5) | ((~ p2(X5) | ! [X8] : (~ p101(X8) | p2(X8) | ~ r1(X5, X8))) & (p2(X5) | ! [X9] : (~ p101(X9) | ~ p2(X9) | ~ r1(X5, X9))))) & (~ p102(X5) | ((~ p3(X5) | ! [X10] : (~ p102(X10) | p3(X10) | ~ r1(X5, X10))) & (p3(X5) | ! [X11] : (~ p102(X11) | ~ p3(X11) | ~ r1(X5, X11))))) & (~ p103(X5) | ((~ p4(X5) | ! [X12] : (~ p103(X12) | p4(X12) | ~ r1(X5, X12))) & (p4(X5) | ! [X13] : (~ p103(X13) | ~ p4(X13) | ~ r1(X5, X13))))) & (~ p104(X5) | ((~ p5(X5) | ! [X14] : (~ p104(X14) | p5(X14) | ~ r1(X5, X14))) & (p5(X5) | ! [X15] : (~ p104(X15) | ~ p5(X15) | ~ r1(X5, X15))))) & (~ p105(X5) | ((~ p6(X5) | ! [X16] : (~ p105(X16) | p6(X16) | ~ r1(X5, X16))) & (p6(X5) | ! [X17] : (~ p105(X17) | ~ p6(X17) | ~ r1(X5, X17))))) & (~ (p100(X5) & ~ p101(X5)) | (~ ! [X18] : (~ (p101(X18) & ~ p102(X18) & p2(X18)) | ~ r1(X5, X18)) & ~ ! [X19] : (~ (p101(X19) & ~ p102(X19) & ~ p2(X19)) | ~ r1(X5, X19)))) & (~ (p101(X5) & ~ p102(X5)) | (~ ! [X20] : (~ (p102(X20) & ~ p103(X20) & p3(X20)) | ~ r1(X5, X20)) & ~ ! [X21] : (~ (p102(X21) & ~ p103(X21) & ~ p3(X21)) | ~ r1(X5, X21)))) & (~ (p102(X5) & ~ p103(X5)) | (~ ! [X22] : (~ (p103(X22) & ~ p104(X22) & p4(X22)) | ~ r1(X5, X22)) & ~ ! [X23] : (~ (p103(X23) & ~ p104(X23) & ~ p4(X23)) | ~ r1(X5, X23)))) & (~ (p103(X5) & ~ p104(X5)) | (~ ! [X24] : (~ (p104(X24) & ~ p105(X24) & p5(X24)) | ~ r1(X5, X24)) & ~ ! [X25] : (~ (p104(X25) & ~ p105(X25) & ~ p5(X25)) | ~ r1(X5, X25)))) & (~ (p104(X5) & ~ p105(X5)) | (~ ! [X26] : (~ (p105(X26) & p6(X26)) | ~ r1(X5, X26)) & ~ ! [X27] : (~ (p105(X27) & ~ p6(X27)) | ~ r1(X5, X27))))) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) | ~ ! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (p3(X32) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X0, X28))), inference(pure_predicate_removal, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (((~ p101(X5) | p100(X5)) & (~ p102(X5) | p101(X5)) & (~ p103(X5) | p102(X5)) & (~ p104(X5) | p103(X5)) & (~ p105(X5) | p104(X5)) & (~ p106(X5) | p105(X5)) & (~ p100(X5) | ((~ p1(X5) | ! [X6] : (~ p100(X6) | p1(X6) | ~ r1(X5, X6))) & (p1(X5) | ! [X7] : (~ p100(X7) | ~ p1(X7) | ~ r1(X5, X7))))) & (~ p101(X5) | ((~ p2(X5) | ! [X8] : (~ p101(X8) | p2(X8) | ~ r1(X5, X8))) & (p2(X5) | ! [X9] : (~ p101(X9) | ~ p2(X9) | ~ r1(X5, X9))))) & (~ p102(X5) | ((~ p3(X5) | ! [X10] : (~ p102(X10) | p3(X10) | ~ r1(X5, X10))) & (p3(X5) | ! [X11] : (~ p102(X11) | ~ p3(X11) | ~ r1(X5, X11))))) & (~ p103(X5) | ((~ p4(X5) | ! [X12] : (~ p103(X12) | p4(X12) | ~ r1(X5, X12))) & (p4(X5) | ! [X13] : (~ p103(X13) | ~ p4(X13) | ~ r1(X5, X13))))) & (~ p104(X5) | ((~ p5(X5) | ! [X14] : (~ p104(X14) | p5(X14) | ~ r1(X5, X14))) & (p5(X5) | ! [X15] : (~ p104(X15) | ~ p5(X15) | ~ r1(X5, X15))))) & (~ p105(X5) | ((~ p6(X5) | ! [X16] : (~ p105(X16) | p6(X16) | ~ r1(X5, X16))) & (p6(X5) | ! [X17] : (~ p105(X17) | ~ p6(X17) | ~ r1(X5, X17))))) & (~ (p100(X5) & ~ p101(X5)) | (~ ! [X18] : (~ (p101(X18) & ~ p102(X18) & p2(X18)) | ~ r1(X5, X18)) & ~ ! [X19] : (~ (p101(X19) & ~ p102(X19) & ~ p2(X19)) | ~ r1(X5, X19)))) & (~ (p101(X5) & ~ p102(X5)) | (~ ! [X20] : (~ (p102(X20) & ~ p103(X20) & p3(X20)) | ~ r1(X5, X20)) & ~ ! [X21] : (~ (p102(X21) & ~ p103(X21) & ~ p3(X21)) | ~ r1(X5, X21)))) & (~ (p102(X5) & ~ p103(X5)) | (~ ! [X22] : (~ (p103(X22) & ~ p104(X22) & p4(X22)) | ~ r1(X5, X22)) & ~ ! [X23] : (~ (p103(X23) & ~ p104(X23) & ~ p4(X23)) | ~ r1(X5, X23)))) & (~ (p103(X5) & ~ p104(X5)) | (~ ! [X24] : (~ (p104(X24) & ~ p105(X24) & p5(X24)) | ~ r1(X5, X24)) & ~ ! [X25] : (~ (p104(X25) & ~ p105(X25) & ~ p5(X25)) | ~ r1(X5, X25)))) & (~ (p104(X5) & ~ p105(X5)) | (~ ! [X26] : (~ (p105(X26) & ~ p106(X26) & p6(X26)) | ~ r1(X5, X26)) & ~ ! [X27] : (~ (p105(X27) & ~ p106(X27) & ~ p6(X27)) | ~ r1(X5, X27))))) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) | ~ ! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (p3(X32) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X0, X28))), inference(flattening, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (((~ p101(X5) | p100(X5)) & (~ p102(X5) | p101(X5)) & (~ p103(X5) | p102(X5)) & (~ p104(X5) | p103(X5)) & (~ p105(X5) | p104(X5)) & (~ p106(X5) | p105(X5)) & (~ p100(X5) | ((~ p1(X5) | ! [X6] : (~ p100(X6) | p1(X6) | ~ r1(X5, X6))) & (p1(X5) | ! [X7] : (~ p100(X7) | ~ p1(X7) | ~ r1(X5, X7))))) & (~ p101(X5) | ((~ p2(X5) | ! [X8] : (~ p101(X8) | p2(X8) | ~ r1(X5, X8))) & (p2(X5) | ! [X9] : (~ p101(X9) | ~ p2(X9) | ~ r1(X5, X9))))) & (~ p102(X5) | ((~ p3(X5) | ! [X10] : (~ p102(X10) | p3(X10) | ~ r1(X5, X10))) & (p3(X5) | ! [X11] : (~ p102(X11) | ~ p3(X11) | ~ r1(X5, X11))))) & (~ p103(X5) | ((~ p4(X5) | ! [X12] : (~ p103(X12) | p4(X12) | ~ r1(X5, X12))) & (p4(X5) | ! [X13] : (~ p103(X13) | ~ p4(X13) | ~ r1(X5, X13))))) & (~ p104(X5) | ((~ p5(X5) | ! [X14] : (~ p104(X14) | p5(X14) | ~ r1(X5, X14))) & (p5(X5) | ! [X15] : (~ p104(X15) | ~ p5(X15) | ~ r1(X5, X15))))) & (~ p105(X5) | ((~ p6(X5) | ! [X16] : (~ p105(X16) | p6(X16) | ~ r1(X5, X16))) & (p6(X5) | ! [X17] : (~ p105(X17) | ~ p6(X17) | ~ r1(X5, X17))))) & (~ (p100(X5) & ~ p101(X5)) | (~ ! [X18] : (~ (p101(X18) & ~ p102(X18) & p2(X18)) | ~ r1(X5, X18)) & ~ ! [X19] : (~ (p101(X19) & ~ p102(X19) & ~ p2(X19)) | ~ r1(X5, X19)))) & (~ (p101(X5) & ~ p102(X5)) | (~ ! [X20] : (~ (p102(X20) & ~ p103(X20) & p3(X20)) | ~ r1(X5, X20)) & ~ ! [X21] : (~ (p102(X21) & ~ p103(X21) & ~ p3(X21)) | ~ r1(X5, X21)))) & (~ (p102(X5) & ~ p103(X5)) | (~ ! [X22] : (~ (p103(X22) & ~ p104(X22) & p4(X22)) | ~ r1(X5, X22)) & ~ ! [X23] : (~ (p103(X23) & ~ p104(X23) & ~ p4(X23)) | ~ r1(X5, X23)))) & (~ (p103(X5) & ~ p104(X5)) | (~ ! [X24] : (~ (p104(X24) & ~ p105(X24) & p5(X24)) | ~ r1(X5, X24)) & ~ ! [X25] : (~ (p104(X25) & ~ p105(X25) & ~ p5(X25)) | ~ r1(X5, X25)))) & (~ (p104(X5) & ~ p105(X5)) | (~ ! [X26] : (~ (p105(X26) & ~ p106(X26) & p6(X26)) | ~ r1(X5, X26)) & ~ ! [X27] : (~ (p105(X27) & ~ p106(X27) & ~ p6(X27)) | ~ r1(X5, X27))))) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) | ~ ! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (p3(X32) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X0, X28))), inference(rectify, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL656+1.005.p', main)).
fof(f1164, plain, (~ p100(sK22) | ~ sP3(sK22) | ~ spl23_121), inference(subsumption_resolution, [], [f1163, f133])).
fof(f133, plain, ~ p101(sK22), inference(cnf_transformation, [], [f63])).
fof(f1163, plain, (p101(sK22) | ~ p100(sK22) | ~ sP3(sK22) | ~ spl23_121), inference(resolution, [], [f1003, f101])).
fof(f101, plain, ! [X0] : (~ p102(sK15(X0)) | p101(X0) | ~ p100(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f45])).
fof(f45, plain, ! [X0] : (~ p100(X0) | p101(X0) | ((p101(sK14(X0)) & ~ p102(sK14(X0)) & p2(sK14(X0)) & r1(X0, sK14(X0))) & (p101(sK15(X0)) & ~ p102(sK15(X0)) & ~ p2(sK15(X0)) & r1(X0, sK15(X0)))) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14, sK15])], [f42, f44, f43])).
fof(f43, plain, ! [X0] : (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) => (p101(sK14(X0)) & ~ p102(sK14(X0)) & p2(sK14(X0)) & r1(X0, sK14(X0)))), introduced(choice_axiom, [])).
fof(f44, plain, ! [X0] : (? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2)) => (p101(sK15(X0)) & ~ p102(sK15(X0)) & ~ p2(sK15(X0)) & r1(X0, sK15(X0)))), introduced(choice_axiom, [])).
fof(f42, plain, ! [X0] : (~ p100(X0) | p101(X0) | (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) & ? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2))) | ~ sP3(X0)), inference(rectify, [], [f41])).
fof(f41, plain, ! [X5] : (~ p100(X5) | p101(X5) | (? [X18] : (p101(X18) & ~ p102(X18) & p2(X18) & r1(X5, X18)) & ? [X19] : (p101(X19) & ~ p102(X19) & ~ p2(X19) & r1(X5, X19))) | ~ sP3(X5)), inference(nnf_transformation, [], [f12])).
fof(f1003, plain, (p102(sK15(sK22)) | ~ spl23_121), inference(avatar_component_clause, [], [f1001])).
fof(f1130, plain, (spl23_120 | ~ spl23_35), inference(avatar_split_clause, [], [f1091, f587, f997])).
fof(f587, plain, (spl23_35 <=> sP11(sK15(sK22))), introduced(avatar_definition, [new_symbols(naming, [spl23_35])])).
fof(f1091, plain, (sP2(sK15(sK22)) | ~ spl23_35), inference(resolution, [], [f589, f68])).
fof(f68, plain, ! [X0] : (~ sP11(X0) | sP2(X0)), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & sP10(X0) & sP9(X0) & sP8(X0) & sP7(X0) & sP6(X0) & sP5(X0) & sP3(X0) & sP2(X0) & sP1(X0) & sP0(X0) & sP4(X0)) | ~ sP11(X0)), inference(rectify, [], [f22])).
fof(f22, plain, ! [X5] : (((~ p101(X5) | p100(X5)) & (~ p102(X5) | p101(X5)) & (~ p103(X5) | p102(X5)) & (~ p104(X5) | p103(X5)) & (~ p105(X5) | p104(X5)) & sP10(X5) & sP9(X5) & sP8(X5) & sP7(X5) & sP6(X5) & sP5(X5) & sP3(X5) & sP2(X5) & sP1(X5) & sP0(X5) & sP4(X5)) | ~ sP11(X5)), inference(nnf_transformation, [], [f20])).
fof(f589, plain, (sP11(sK15(sK22)) | ~ spl23_35), inference(avatar_component_clause, [], [f587])).
fof(f1079, plain, (~ spl23_6 | spl23_117), inference(avatar_contradiction_clause, [], [f1078])).
fof(f1078, plain, ($false | (~ spl23_6 | spl23_117)), inference(subsumption_resolution, [], [f1077, f334])).
fof(f1077, plain, (~ sP3(sK22) | spl23_117), inference(subsumption_resolution, [], [f1076, f134])).
fof(f1076, plain, (~ p100(sK22) | ~ sP3(sK22) | spl23_117), inference(subsumption_resolution, [], [f1075, f133])).
fof(f1075, plain, (p101(sK22) | ~ p100(sK22) | ~ sP3(sK22) | spl23_117), inference(resolution, [], [f986, f102])).
fof(f102, plain, ! [X0] : (p101(sK15(X0)) | p101(X0) | ~ p100(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f45])).
fof(f986, plain, (~ p101(sK15(sK22)) | spl23_117), inference(avatar_component_clause, [], [f985])).
fof(f641, plain, (~ spl23_6 | spl23_43), inference(avatar_split_clause, [], [f637, f639, f333])).
fof(f637, plain, ! [X4] : (~ r1(sK15(sK22), X4) | p3(X4) | ~ sP3(sK22)), inference(subsumption_resolution, [], [f636, f134])).
fof(f636, plain, ! [X4] : (~ r1(sK15(sK22), X4) | p3(X4) | ~ p100(sK22) | ~ sP3(sK22)), inference(subsumption_resolution, [], [f426, f133])).
fof(f426, plain, ! [X4] : (~ r1(sK15(sK22), X4) | p3(X4) | p101(sK22) | ~ p100(sK22) | ~ sP3(sK22)), inference(resolution, [], [f400, f99])).
fof(f99, plain, ! [X0] : (r1(X0, sK15(X0)) | p101(X0) | ~ p100(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f45])).
fof(f400, plain, ! [X0, X1] : (~ r1(sK22, X0) | ~ r1(X0, X1) | p3(X1)), inference(resolution, [], [f389, f64])).
fof(f64, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL656+1.005.p', reflexivity)).
fof(f389, plain, ! [X2, X0, X1] : (~ r1(sK22, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | p3(X2)), inference(resolution, [], [f299, f64])).
fof(f299, plain, ! [X2, X0, X3, X1] : (~ r1(sK22, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | p3(X1)), inference(resolution, [], [f131, f64])).
fof(f131, plain, ! [X6, X10, X8, X7, X9] : (~ r1(sK22, X6) | ~ r1(X9, X10) | ~ r1(X8, X9) | ~ r1(X7, X8) | ~ r1(X6, X7) | p3(X10)), inference(cnf_transformation, [], [f63])).
fof(f590, plain, (~ spl23_6 | spl23_35), inference(avatar_split_clause, [], [f585, f587, f333])).
fof(f585, plain, (sP11(sK15(sK22)) | ~ sP3(sK22)), inference(subsumption_resolution, [], [f584, f134])).
fof(f584, plain, (sP11(sK15(sK22)) | ~ p100(sK22) | ~ sP3(sK22)), inference(subsumption_resolution, [], [f470, f133])).
fof(f470, plain, (sP11(sK15(sK22)) | p101(sK22) | ~ p100(sK22) | ~ sP3(sK22)), inference(resolution, [], [f455, f99])).
fof(f455, plain, ! [X0] : (~ r1(sK22, X0) | sP11(X0)), inference(resolution, [], [f444, f64])).
fof(f444, plain, ! [X0, X1] : (~ r1(sK22, X0) | ~ r1(X0, X1) | sP11(X1)), inference(resolution, [], [f411, f64])).
fof(f411, plain, ! [X2, X0, X1] : (~ r1(sK22, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP11(X2)), inference(resolution, [], [f378, f64])).
fof(f378, plain, ! [X2, X0, X3, X1] : (~ r1(sK22, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | sP11(X1)), inference(resolution, [], [f132, f64])).
fof(f132, plain, ! [X4, X2, X5, X3, X1] : (~ r1(sK22, X1) | ~ r1(X4, X5) | ~ r1(X3, X4) | ~ r1(X2, X3) | ~ r1(X1, X2) | sP11(X5)), inference(cnf_transformation, [], [f63])).
fof(f583, plain, spl23_6, inference(avatar_split_clause, [], [f481, f333])).
fof(f481, plain, sP3(sK22), inference(resolution, [], [f466, f69])).
fof(f69, plain, ! [X0] : (~ sP11(X0) | sP3(X0)), inference(cnf_transformation, [], [f23])).
fof(f466, plain, sP11(sK22), inference(resolution, [], [f455, f64])).