fof(f16563, plain, $false, inference(avatar_sat_refutation, [], [f423, f441, f460, f480, f501, f522, f575, f581, f583, f586, f590, f592, f595, f599, f604, f606, f609, f613, f618, f801, f1473, f2290, f2312, f2366, f2960, f3029, f3074, f3272, f3344, f3362, f3495, f3592, f3628, f3646, f4299, f4508, f6407, f6496, f6989, f7270, f7988, f8304, f8308, f8322, f8334, f8336, f8998, f9208, f9216, f9220, f9513, f9515, f9925, f9937, f10133, f10819, f11235, f11655, f11676, f11684, f11698, f11706, f11712, f11795, f11812, f11818, f12412, f12515, f12516, f12520, f12588, f12600, f12601, f12808, f12810, f12904, f12999, f13076, f13077, f13078, f13080, f13243, f13249, f13277, f13640, f13645, f13657, f13857, f13873, f13888, f14164, f14173, f14213, f14315, f14716, f14723, f14733, f14738, f14757, f14758, f14760, f14773, f14780, f14782, f14884, f14908, f14916, f14919, f14923, f14931, f15346, f15429, f15436, f15441, f15521, f15522, f15524, f15616, f15617, f15695, f15856, f15877, f15879, f16038, f16144, f16205, f16241, f16267, f16357, f16460, f16462, f16470, f16472, f16510, f16562])).
fof(f16562, plain, (~ spl92_21 | ~ spl92_1411), inference(avatar_contradiction_clause, [], [f16561])).
fof(f16561, plain, ($false | (~ spl92_21 | ~ spl92_1411)), inference(subsumption_resolution, [], [f16560, f527])).
fof(f527, plain, sP13(sK91), inference(resolution, [], [f523, f229])).
fof(f229, plain, ! [X0] : (~ sP40(X0) | sP13(X0)), inference(cnf_transformation, [], [f51])).
fof(f51, plain, ! [X0] : (((~ p101(X0) | ~ p201(X0)) & (~ p101(X0) | ~ p301(X0)) & (~ p101(X0) | ~ p401(X0)) & (~ p101(X0) | ~ p501(X0)) & (~ p101(X0) | ~ p601(X0)) & (~ p201(X0) | ~ p301(X0)) & (~ p201(X0) | ~ p401(X0)) & (~ p201(X0) | ~ p501(X0)) & (~ p201(X0) | ~ p601(X0)) & (~ p301(X0) | ~ p401(X0)) & (~ p301(X0) | ~ p501(X0)) & (~ p301(X0) | ~ p601(X0)) & (~ p401(X0) | ~ p501(X0)) & (~ p401(X0) | ~ p601(X0)) & (~ p501(X0) | ~ p601(X0)) & sP39(X0) & sP38(X0) & sP37(X0) & sP36(X0) & sP35(X0) & (~ p202(X0) | ~ p302(X0)) & (~ p202(X0) | ~ p402(X0)) & (~ p202(X0) | ~ p502(X0)) & (~ p202(X0) | ~ p602(X0)) & (~ p302(X0) | ~ p402(X0)) & (~ p302(X0) | ~ p502(X0)) & (~ p302(X0) | ~ p602(X0)) & (~ p402(X0) | ~ p502(X0)) & (~ p402(X0) | ~ p602(X0)) & (~ p502(X0) | ~ p602(X0)) & sP9(X0) & sP34(X0) & sP33(X0) & sP32(X0) & sP31(X0) & sP30(X0) & sP29(X0) & sP28(X0) & sP27(X0) & (~ p303(X0) | ~ p403(X0)) & (~ p303(X0) | ~ p503(X0)) & (~ p303(X0) | ~ p603(X0)) & (~ p403(X0) | ~ p503(X0)) & (~ p403(X0) | ~ p603(X0)) & (~ p503(X0) | ~ p603(X0)) & sP8(X0) & sP7(X0) & sP26(X0) & sP25(X0) & sP24(X0) & sP6(X0) & sP23(X0) & sP22(X0) & sP21(X0) & sP20(X0) & sP19(X0) & sP18(X0) & (~ p404(X0) | ~ p504(X0)) & (~ p404(X0) | ~ p604(X0)) & (~ p504(X0) | ~ p604(X0)) & sP5(X0) & sP4(X0) & sP3(X0) & sP17(X0) & sP16(X0) & sP2(X0) & sP1(X0) & sP15(X0) & sP14(X0) & sP0(X0) & sP13(X0) & sP12(X0) & sP11(X0) & sP10(X0) & (~ p505(X0) | ~ p605(X0))) | ~ sP40(X0)), inference(rectify, [], [f50])).
fof(f50, plain, ! [X11] : (((~ p101(X11) | ~ p201(X11)) & (~ p101(X11) | ~ p301(X11)) & (~ p101(X11) | ~ p401(X11)) & (~ p101(X11) | ~ p501(X11)) & (~ p101(X11) | ~ p601(X11)) & (~ p201(X11) | ~ p301(X11)) & (~ p201(X11) | ~ p401(X11)) & (~ p201(X11) | ~ p501(X11)) & (~ p201(X11) | ~ p601(X11)) & (~ p301(X11) | ~ p401(X11)) & (~ p301(X11) | ~ p501(X11)) & (~ p301(X11) | ~ p601(X11)) & (~ p401(X11) | ~ p501(X11)) & (~ p401(X11) | ~ p601(X11)) & (~ p501(X11) | ~ p601(X11)) & sP39(X11) & sP38(X11) & sP37(X11) & sP36(X11) & sP35(X11) & (~ p202(X11) | ~ p302(X11)) & (~ p202(X11) | ~ p402(X11)) & (~ p202(X11) | ~ p502(X11)) & (~ p202(X11) | ~ p602(X11)) & (~ p302(X11) | ~ p402(X11)) & (~ p302(X11) | ~ p502(X11)) & (~ p302(X11) | ~ p602(X11)) & (~ p402(X11) | ~ p502(X11)) & (~ p402(X11) | ~ p602(X11)) & (~ p502(X11) | ~ p602(X11)) & sP9(X11) & sP34(X11) & sP33(X11) & sP32(X11) & sP31(X11) & sP30(X11) & sP29(X11) & sP28(X11) & sP27(X11) & (~ p303(X11) | ~ p403(X11)) & (~ p303(X11) | ~ p503(X11)) & (~ p303(X11) | ~ p603(X11)) & (~ p403(X11) | ~ p503(X11)) & (~ p403(X11) | ~ p603(X11)) & (~ p503(X11) | ~ p603(X11)) & sP8(X11) & sP7(X11) & sP26(X11) & sP25(X11) & sP24(X11) & sP6(X11) & sP23(X11) & sP22(X11) & sP21(X11) & sP20(X11) & sP19(X11) & sP18(X11) & (~ p404(X11) | ~ p504(X11)) & (~ p404(X11) | ~ p604(X11)) & (~ p504(X11) | ~ p604(X11)) & sP5(X11) & sP4(X11) & sP3(X11) & sP17(X11) & sP16(X11) & sP2(X11) & sP1(X11) & sP15(X11) & sP14(X11) & sP0(X11) & sP13(X11) & sP12(X11) & sP11(X11) & sP10(X11) & (~ p505(X11) | ~ p605(X11))) | ~ sP40(X11)), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ! [X11] : (((~ p101(X11) | ~ p201(X11)) & (~ p101(X11) | ~ p301(X11)) & (~ p101(X11) | ~ p401(X11)) & (~ p101(X11) | ~ p501(X11)) & (~ p101(X11) | ~ p601(X11)) & (~ p201(X11) | ~ p301(X11)) & (~ p201(X11) | ~ p401(X11)) & (~ p201(X11) | ~ p501(X11)) & (~ p201(X11) | ~ p601(X11)) & (~ p301(X11) | ~ p401(X11)) & (~ p301(X11) | ~ p501(X11)) & (~ p301(X11) | ~ p601(X11)) & (~ p401(X11) | ~ p501(X11)) & (~ p401(X11) | ~ p601(X11)) & (~ p501(X11) | ~ p601(X11)) & sP39(X11) & sP38(X11) & sP37(X11) & sP36(X11) & sP35(X11) & (~ p202(X11) | ~ p302(X11)) & (~ p202(X11) | ~ p402(X11)) & (~ p202(X11) | ~ p502(X11)) & (~ p202(X11) | ~ p602(X11)) & (~ p302(X11) | ~ p402(X11)) & (~ p302(X11) | ~ p502(X11)) & (~ p302(X11) | ~ p602(X11)) & (~ p402(X11) | ~ p502(X11)) & (~ p402(X11) | ~ p602(X11)) & (~ p502(X11) | ~ p602(X11)) & sP9(X11) & sP34(X11) & sP33(X11) & sP32(X11) & sP31(X11) & sP30(X11) & sP29(X11) & sP28(X11) & sP27(X11) & (~ p303(X11) | ~ p403(X11)) & (~ p303(X11) | ~ p503(X11)) & (~ p303(X11) | ~ p603(X11)) & (~ p403(X11) | ~ p503(X11)) & (~ p403(X11) | ~ p603(X11)) & (~ p503(X11) | ~ p603(X11)) & sP8(X11) & sP7(X11) & sP26(X11) & sP25(X11) & sP24(X11) & sP6(X11) & sP23(X11) & sP22(X11) & sP21(X11) & sP20(X11) & sP19(X11) & sP18(X11) & (~ p404(X11) | ~ p504(X11)) & (~ p404(X11) | ~ p604(X11)) & (~ p504(X11) | ~ p604(X11)) & sP5(X11) & sP4(X11) & sP3(X11) & sP17(X11) & sP16(X11) & sP2(X11) & sP1(X11) & sP15(X11) & sP14(X11) & sP0(X11) & sP13(X11) & sP12(X11) & sP11(X11) & sP10(X11) & (~ p505(X11) | ~ p605(X11))) | ~ sP40(X11)), inference(usedef, [], [e48])).
fof(e48, plain, ! [X11] : (sP40(X11) <=> ((~ p101(X11) | ~ p201(X11)) & (~ p101(X11) | ~ p301(X11)) & (~ p101(X11) | ~ p401(X11)) & (~ p101(X11) | ~ p501(X11)) & (~ p101(X11) | ~ p601(X11)) & (~ p201(X11) | ~ p301(X11)) & (~ p201(X11) | ~ p401(X11)) & (~ p201(X11) | ~ p501(X11)) & (~ p201(X11) | ~ p601(X11)) & (~ p301(X11) | ~ p401(X11)) & (~ p301(X11) | ~ p501(X11)) & (~ p301(X11) | ~ p601(X11)) & (~ p401(X11) | ~ p501(X11)) & (~ p401(X11) | ~ p601(X11)) & (~ p501(X11) | ~ p601(X11)) & sP39(X11) & sP38(X11) & sP37(X11) & sP36(X11) & sP35(X11) & (~ p202(X11) | ~ p302(X11)) & (~ p202(X11) | ~ p402(X11)) & (~ p202(X11) | ~ p502(X11)) & (~ p202(X11) | ~ p602(X11)) & (~ p302(X11) | ~ p402(X11)) & (~ p302(X11) | ~ p502(X11)) & (~ p302(X11) | ~ p602(X11)) & (~ p402(X11) | ~ p502(X11)) & (~ p402(X11) | ~ p602(X11)) & (~ p502(X11) | ~ p602(X11)) & sP9(X11) & sP34(X11) & sP33(X11) & sP32(X11) & sP31(X11) & sP30(X11) & sP29(X11) & sP28(X11) & sP27(X11) & (~ p303(X11) | ~ p403(X11)) & (~ p303(X11) | ~ p503(X11)) & (~ p303(X11) | ~ p603(X11)) & (~ p403(X11) | ~ p503(X11)) & (~ p403(X11) | ~ p603(X11)) & (~ p503(X11) | ~ p603(X11)) & sP8(X11) & sP7(X11) & sP26(X11) & sP25(X11) & sP24(X11) & sP6(X11) & sP23(X11) & sP22(X11) & sP21(X11) & sP20(X11) & sP19(X11) & sP18(X11) & (~ p404(X11) | ~ p504(X11)) & (~ p404(X11) | ~ p604(X11)) & (~ p504(X11) | ~ p604(X11)) & sP5(X11) & sP4(X11) & sP3(X11) & sP17(X11) & sP16(X11) & sP2(X11) & sP1(X11) & sP15(X11) & sP14(X11) & sP0(X11) & sP13(X11) & sP12(X11) & sP11(X11) & sP10(X11) & (~ p505(X11) | ~ p605(X11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f523, plain, sP40(sK91), inference(resolution, [], [f400, f224])).
fof(f224, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL666+1.005.p', reflexivity)).
fof(f400, plain, ! [X11] : (~ r1(sK91, X11) | sP40(X11)), inference(cnf_transformation, [], [f223])).
fof(f223, plain, ((p101(sK91) | ! [X1] : (p102(X1) | ~ r1(sK91, X1)) | ! [X2] : (p103(X2) | ~ r1(sK91, X2)) | ! [X3] : (p104(X3) | ~ r1(sK91, X3)) | ! [X4] : (p105(X4) | ~ r1(sK91, X4))) & (p201(sK91) | p202(sK91) | ! [X5] : (p203(X5) | ~ r1(sK91, X5)) | ! [X6] : (p204(X6) | ~ r1(sK91, X6)) | ! [X7] : (p205(X7) | ~ r1(sK91, X7))) & (p301(sK91) | p302(sK91) | p303(sK91) | ! [X8] : (p304(X8) | ~ r1(sK91, X8)) | ! [X9] : (p305(X9) | ~ r1(sK91, X9))) & (p401(sK91) | p402(sK91) | p403(sK91) | p404(sK91) | ! [X10] : (p405(X10) | ~ r1(sK91, X10))) & (p501(sK91) | p502(sK91) | p503(sK91) | p504(sK91) | p505(sK91)) & (p601(sK91) | p602(sK91) | p603(sK91) | p604(sK91) | p605(sK91)) & ! [X11] : (sP40(X11) | ~ r1(sK91, X11))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK91])], [f49, f222])).
fof(f222, plain, (? [X0] : ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0)) & ! [X11] : (sP40(X11) | ~ r1(X0, X11))) => ((p101(sK91) | ! [X1] : (p102(X1) | ~ r1(sK91, X1)) | ! [X2] : (p103(X2) | ~ r1(sK91, X2)) | ! [X3] : (p104(X3) | ~ r1(sK91, X3)) | ! [X4] : (p105(X4) | ~ r1(sK91, X4))) & (p201(sK91) | p202(sK91) | ! [X5] : (p203(X5) | ~ r1(sK91, X5)) | ! [X6] : (p204(X6) | ~ r1(sK91, X6)) | ! [X7] : (p205(X7) | ~ r1(sK91, X7))) & (p301(sK91) | p302(sK91) | p303(sK91) | ! [X8] : (p304(X8) | ~ r1(sK91, X8)) | ! [X9] : (p305(X9) | ~ r1(sK91, X9))) & (p401(sK91) | p402(sK91) | p403(sK91) | p404(sK91) | ! [X10] : (p405(X10) | ~ r1(sK91, X10))) & (p501(sK91) | p502(sK91) | p503(sK91) | p504(sK91) | p505(sK91)) & (p601(sK91) | p602(sK91) | p603(sK91) | p604(sK91) | p605(sK91)) & ! [X11] : (sP40(X11) | ~ r1(sK91, X11)))), introduced(choice_axiom, [])).
fof(f49, plain, ? [X0] : ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0)) & ! [X11] : (sP40(X11) | ~ r1(X0, X11))), inference(definition_folding, [], [f7, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8])).
fof(f8, plain, ! [X11] : (? [X56] : (~ p305(X56) & r1(X11, X56)) | ? [X57] : (~ p405(X57) & r1(X11, X57)) | ~ sP0(X11)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X11] : (sP0(X11) <=> (? [X56] : (~ p305(X56) & r1(X11, X56)) | ? [X57] : (~ p405(X57) & r1(X11, X57)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f9, plain, ! [X11] : (? [X52] : (~ p205(X52) & r1(X11, X52)) | ? [X53] : (~ p405(X53) & r1(X11, X53)) | ~ sP1(X11)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X11] : (sP1(X11) <=> (? [X52] : (~ p205(X52) & r1(X11, X52)) | ? [X53] : (~ p405(X53) & r1(X11, X53)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f10, plain, ! [X11] : (? [X50] : (~ p205(X50) & r1(X11, X50)) | ? [X51] : (~ p305(X51) & r1(X11, X51)) | ~ sP2(X11)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X11] : (sP2(X11) <=> (? [X50] : (~ p205(X50) & r1(X11, X50)) | ? [X51] : (~ p305(X51) & r1(X11, X51)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f11, plain, ! [X11] : (? [X46] : (~ p105(X46) & r1(X11, X46)) | ? [X47] : (~ p405(X47) & r1(X11, X47)) | ~ sP3(X11)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X11] : (sP3(X11) <=> (? [X46] : (~ p105(X46) & r1(X11, X46)) | ? [X47] : (~ p405(X47) & r1(X11, X47)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f12, plain, ! [X11] : (? [X44] : (~ p105(X44) & r1(X11, X44)) | ? [X45] : (~ p305(X45) & r1(X11, X45)) | ~ sP4(X11)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X11] : (sP4(X11) <=> (? [X44] : (~ p105(X44) & r1(X11, X44)) | ? [X45] : (~ p305(X45) & r1(X11, X45)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f13, plain, ! [X11] : (? [X42] : (~ p105(X42) & r1(X11, X42)) | ? [X43] : (~ p205(X43) & r1(X11, X43)) | ~ sP5(X11)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X11] : (sP5(X11) <=> (? [X42] : (~ p105(X42) & r1(X11, X42)) | ? [X43] : (~ p205(X43) & r1(X11, X43)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f14, plain, ! [X11] : (? [X34] : (~ p204(X34) & r1(X11, X34)) | ? [X35] : (~ p304(X35) & r1(X11, X35)) | ~ sP6(X11)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X11] : (sP6(X11) <=> (? [X34] : (~ p204(X34) & r1(X11, X34)) | ? [X35] : (~ p304(X35) & r1(X11, X35)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f15, plain, ! [X11] : (? [X29] : (~ p104(X29) & r1(X11, X29)) | ? [X30] : (~ p304(X30) & r1(X11, X30)) | ~ sP7(X11)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X11] : (sP7(X11) <=> (? [X29] : (~ p104(X29) & r1(X11, X29)) | ? [X30] : (~ p304(X30) & r1(X11, X30)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f16, plain, ! [X11] : (? [X27] : (~ p104(X27) & r1(X11, X27)) | ? [X28] : (~ p204(X28) & r1(X11, X28)) | ~ sP8(X11)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X11] : (sP8(X11) <=> (? [X27] : (~ p104(X27) & r1(X11, X27)) | ? [X28] : (~ p204(X28) & r1(X11, X28)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f17, plain, ! [X11] : (? [X17] : (~ p103(X17) & r1(X11, X17)) | ? [X18] : (~ p203(X18) & r1(X11, X18)) | ~ sP9(X11)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X11] : (sP9(X11) <=> (? [X17] : (~ p103(X17) & r1(X11, X17)) | ? [X18] : (~ p203(X18) & r1(X11, X18)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f18, plain, ! [X11] : (? [X61] : (~ p405(X61) & r1(X11, X61)) | ~ p605(X11) | ~ sP10(X11)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X11] : (sP10(X11) <=> (? [X61] : (~ p405(X61) & r1(X11, X61)) | ~ p605(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f19, plain, ! [X11] : (? [X60] : (~ p405(X60) & r1(X11, X60)) | ~ p505(X11) | ~ sP11(X11)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X11] : (sP11(X11) <=> (? [X60] : (~ p405(X60) & r1(X11, X60)) | ~ p505(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f20, plain, ! [X11] : (? [X59] : (~ p305(X59) & r1(X11, X59)) | ~ p605(X11) | ~ sP12(X11)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X11] : (sP12(X11) <=> (? [X59] : (~ p305(X59) & r1(X11, X59)) | ~ p605(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f21, plain, ! [X11] : (? [X58] : (~ p305(X58) & r1(X11, X58)) | ~ p505(X11) | ~ sP13(X11)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X11] : (sP13(X11) <=> (? [X58] : (~ p305(X58) & r1(X11, X58)) | ~ p505(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f22, plain, ! [X11] : (? [X55] : (~ p205(X55) & r1(X11, X55)) | ~ p605(X11) | ~ sP14(X11)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X11] : (sP14(X11) <=> (? [X55] : (~ p205(X55) & r1(X11, X55)) | ~ p605(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f23, plain, ! [X11] : (? [X54] : (~ p205(X54) & r1(X11, X54)) | ~ p505(X11) | ~ sP15(X11)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X11] : (sP15(X11) <=> (? [X54] : (~ p205(X54) & r1(X11, X54)) | ~ p505(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f24, plain, ! [X11] : (? [X49] : (~ p105(X49) & r1(X11, X49)) | ~ p605(X11) | ~ sP16(X11)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X11] : (sP16(X11) <=> (? [X49] : (~ p105(X49) & r1(X11, X49)) | ~ p605(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f25, plain, ! [X11] : (? [X48] : (~ p105(X48) & r1(X11, X48)) | ~ p505(X11) | ~ sP17(X11)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X11] : (sP17(X11) <=> (? [X48] : (~ p105(X48) & r1(X11, X48)) | ~ p505(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f26, plain, ! [X11] : (? [X41] : (~ p304(X41) & r1(X11, X41)) | ~ p604(X11) | ~ sP18(X11)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X11] : (sP18(X11) <=> (? [X41] : (~ p304(X41) & r1(X11, X41)) | ~ p604(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f27, plain, ! [X11] : (? [X40] : (~ p304(X40) & r1(X11, X40)) | ~ p504(X11) | ~ sP19(X11)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X11] : (sP19(X11) <=> (? [X40] : (~ p304(X40) & r1(X11, X40)) | ~ p504(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f28, plain, ! [X11] : (? [X39] : (~ p304(X39) & r1(X11, X39)) | ~ p404(X11) | ~ sP20(X11)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X11] : (sP20(X11) <=> (? [X39] : (~ p304(X39) & r1(X11, X39)) | ~ p404(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f29, plain, ! [X11] : (? [X38] : (~ p204(X38) & r1(X11, X38)) | ~ p604(X11) | ~ sP21(X11)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X11] : (sP21(X11) <=> (? [X38] : (~ p204(X38) & r1(X11, X38)) | ~ p604(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f30, plain, ! [X11] : (? [X37] : (~ p204(X37) & r1(X11, X37)) | ~ p504(X11) | ~ sP22(X11)), inference(usedef, [], [e30])).
fof(e30, plain, ! [X11] : (sP22(X11) <=> (? [X37] : (~ p204(X37) & r1(X11, X37)) | ~ p504(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f31, plain, ! [X11] : (? [X36] : (~ p204(X36) & r1(X11, X36)) | ~ p404(X11) | ~ sP23(X11)), inference(usedef, [], [e31])).
fof(e31, plain, ! [X11] : (sP23(X11) <=> (? [X36] : (~ p204(X36) & r1(X11, X36)) | ~ p404(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f32, plain, ! [X11] : (? [X33] : (~ p104(X33) & r1(X11, X33)) | ~ p604(X11) | ~ sP24(X11)), inference(usedef, [], [e32])).
fof(e32, plain, ! [X11] : (sP24(X11) <=> (? [X33] : (~ p104(X33) & r1(X11, X33)) | ~ p604(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f33, plain, ! [X11] : (? [X32] : (~ p104(X32) & r1(X11, X32)) | ~ p504(X11) | ~ sP25(X11)), inference(usedef, [], [e33])).
fof(e33, plain, ! [X11] : (sP25(X11) <=> (? [X32] : (~ p104(X32) & r1(X11, X32)) | ~ p504(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f34, plain, ! [X11] : (? [X31] : (~ p104(X31) & r1(X11, X31)) | ~ p404(X11) | ~ sP26(X11)), inference(usedef, [], [e34])).
fof(e34, plain, ! [X11] : (sP26(X11) <=> (? [X31] : (~ p104(X31) & r1(X11, X31)) | ~ p404(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f35, plain, ! [X11] : (? [X26] : (~ p203(X26) & r1(X11, X26)) | ~ p603(X11) | ~ sP27(X11)), inference(usedef, [], [e35])).
fof(e35, plain, ! [X11] : (sP27(X11) <=> (? [X26] : (~ p203(X26) & r1(X11, X26)) | ~ p603(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f36, plain, ! [X11] : (? [X25] : (~ p203(X25) & r1(X11, X25)) | ~ p503(X11) | ~ sP28(X11)), inference(usedef, [], [e36])).
fof(e36, plain, ! [X11] : (sP28(X11) <=> (? [X25] : (~ p203(X25) & r1(X11, X25)) | ~ p503(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f37, plain, ! [X11] : (? [X24] : (~ p203(X24) & r1(X11, X24)) | ~ p403(X11) | ~ sP29(X11)), inference(usedef, [], [e37])).
fof(e37, plain, ! [X11] : (sP29(X11) <=> (? [X24] : (~ p203(X24) & r1(X11, X24)) | ~ p403(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f38, plain, ! [X11] : (? [X23] : (~ p203(X23) & r1(X11, X23)) | ~ p303(X11) | ~ sP30(X11)), inference(usedef, [], [e38])).
fof(e38, plain, ! [X11] : (sP30(X11) <=> (? [X23] : (~ p203(X23) & r1(X11, X23)) | ~ p303(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f39, plain, ! [X11] : (? [X22] : (~ p103(X22) & r1(X11, X22)) | ~ p603(X11) | ~ sP31(X11)), inference(usedef, [], [e39])).
fof(e39, plain, ! [X11] : (sP31(X11) <=> (? [X22] : (~ p103(X22) & r1(X11, X22)) | ~ p603(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f40, plain, ! [X11] : (? [X21] : (~ p103(X21) & r1(X11, X21)) | ~ p503(X11) | ~ sP32(X11)), inference(usedef, [], [e40])).
fof(e40, plain, ! [X11] : (sP32(X11) <=> (? [X21] : (~ p103(X21) & r1(X11, X21)) | ~ p503(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f41, plain, ! [X11] : (? [X20] : (~ p103(X20) & r1(X11, X20)) | ~ p403(X11) | ~ sP33(X11)), inference(usedef, [], [e41])).
fof(e41, plain, ! [X11] : (sP33(X11) <=> (? [X20] : (~ p103(X20) & r1(X11, X20)) | ~ p403(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f42, plain, ! [X11] : (? [X19] : (~ p103(X19) & r1(X11, X19)) | ~ p303(X11) | ~ sP34(X11)), inference(usedef, [], [e42])).
fof(e42, plain, ! [X11] : (sP34(X11) <=> (? [X19] : (~ p103(X19) & r1(X11, X19)) | ~ p303(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f43, plain, ! [X11] : (? [X16] : (~ p102(X16) & r1(X11, X16)) | ~ p602(X11) | ~ sP35(X11)), inference(usedef, [], [e43])).
fof(e43, plain, ! [X11] : (sP35(X11) <=> (? [X16] : (~ p102(X16) & r1(X11, X16)) | ~ p602(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f44, plain, ! [X11] : (? [X15] : (~ p102(X15) & r1(X11, X15)) | ~ p502(X11) | ~ sP36(X11)), inference(usedef, [], [e44])).
fof(e44, plain, ! [X11] : (sP36(X11) <=> (? [X15] : (~ p102(X15) & r1(X11, X15)) | ~ p502(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f45, plain, ! [X11] : (? [X14] : (~ p102(X14) & r1(X11, X14)) | ~ p402(X11) | ~ sP37(X11)), inference(usedef, [], [e45])).
fof(e45, plain, ! [X11] : (sP37(X11) <=> (? [X14] : (~ p102(X14) & r1(X11, X14)) | ~ p402(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f46, plain, ! [X11] : (? [X13] : (~ p102(X13) & r1(X11, X13)) | ~ p302(X11) | ~ sP38(X11)), inference(usedef, [], [e46])).
fof(e46, plain, ! [X11] : (sP38(X11) <=> (? [X13] : (~ p102(X13) & r1(X11, X13)) | ~ p302(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f47, plain, ! [X11] : (? [X12] : (~ p102(X12) & r1(X11, X12)) | ~ p202(X11) | ~ sP39(X11)), inference(usedef, [], [e47])).
fof(e47, plain, ! [X11] : (sP39(X11) <=> (? [X12] : (~ p102(X12) & r1(X11, X12)) | ~ p202(X11))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f7, plain, ? [X0] : ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0)) & ! [X11] : (((~ p101(X11) | ~ p201(X11)) & (~ p101(X11) | ~ p301(X11)) & (~ p101(X11) | ~ p401(X11)) & (~ p101(X11) | ~ p501(X11)) & (~ p101(X11) | ~ p601(X11)) & (~ p201(X11) | ~ p301(X11)) & (~ p201(X11) | ~ p401(X11)) & (~ p201(X11) | ~ p501(X11)) & (~ p201(X11) | ~ p601(X11)) & (~ p301(X11) | ~ p401(X11)) & (~ p301(X11) | ~ p501(X11)) & (~ p301(X11) | ~ p601(X11)) & (~ p401(X11) | ~ p501(X11)) & (~ p401(X11) | ~ p601(X11)) & (~ p501(X11) | ~ p601(X11)) & (? [X12] : (~ p102(X12) & r1(X11, X12)) | ~ p202(X11)) & (? [X13] : (~ p102(X13) & r1(X11, X13)) | ~ p302(X11)) & (? [X14] : (~ p102(X14) & r1(X11, X14)) | ~ p402(X11)) & (? [X15] : (~ p102(X15) & r1(X11, X15)) | ~ p502(X11)) & (? [X16] : (~ p102(X16) & r1(X11, X16)) | ~ p602(X11)) & (~ p202(X11) | ~ p302(X11)) & (~ p202(X11) | ~ p402(X11)) & (~ p202(X11) | ~ p502(X11)) & (~ p202(X11) | ~ p602(X11)) & (~ p302(X11) | ~ p402(X11)) & (~ p302(X11) | ~ p502(X11)) & (~ p302(X11) | ~ p602(X11)) & (~ p402(X11) | ~ p502(X11)) & (~ p402(X11) | ~ p602(X11)) & (~ p502(X11) | ~ p602(X11)) & (? [X17] : (~ p103(X17) & r1(X11, X17)) | ? [X18] : (~ p203(X18) & r1(X11, X18))) & (? [X19] : (~ p103(X19) & r1(X11, X19)) | ~ p303(X11)) & (? [X20] : (~ p103(X20) & r1(X11, X20)) | ~ p403(X11)) & (? [X21] : (~ p103(X21) & r1(X11, X21)) | ~ p503(X11)) & (? [X22] : (~ p103(X22) & r1(X11, X22)) | ~ p603(X11)) & (? [X23] : (~ p203(X23) & r1(X11, X23)) | ~ p303(X11)) & (? [X24] : (~ p203(X24) & r1(X11, X24)) | ~ p403(X11)) & (? [X25] : (~ p203(X25) & r1(X11, X25)) | ~ p503(X11)) & (? [X26] : (~ p203(X26) & r1(X11, X26)) | ~ p603(X11)) & (~ p303(X11) | ~ p403(X11)) & (~ p303(X11) | ~ p503(X11)) & (~ p303(X11) | ~ p603(X11)) & (~ p403(X11) | ~ p503(X11)) & (~ p403(X11) | ~ p603(X11)) & (~ p503(X11) | ~ p603(X11)) & (? [X27] : (~ p104(X27) & r1(X11, X27)) | ? [X28] : (~ p204(X28) & r1(X11, X28))) & (? [X29] : (~ p104(X29) & r1(X11, X29)) | ? [X30] : (~ p304(X30) & r1(X11, X30))) & (? [X31] : (~ p104(X31) & r1(X11, X31)) | ~ p404(X11)) & (? [X32] : (~ p104(X32) & r1(X11, X32)) | ~ p504(X11)) & (? [X33] : (~ p104(X33) & r1(X11, X33)) | ~ p604(X11)) & (? [X34] : (~ p204(X34) & r1(X11, X34)) | ? [X35] : (~ p304(X35) & r1(X11, X35))) & (? [X36] : (~ p204(X36) & r1(X11, X36)) | ~ p404(X11)) & (? [X37] : (~ p204(X37) & r1(X11, X37)) | ~ p504(X11)) & (? [X38] : (~ p204(X38) & r1(X11, X38)) | ~ p604(X11)) & (? [X39] : (~ p304(X39) & r1(X11, X39)) | ~ p404(X11)) & (? [X40] : (~ p304(X40) & r1(X11, X40)) | ~ p504(X11)) & (? [X41] : (~ p304(X41) & r1(X11, X41)) | ~ p604(X11)) & (~ p404(X11) | ~ p504(X11)) & (~ p404(X11) | ~ p604(X11)) & (~ p504(X11) | ~ p604(X11)) & (? [X42] : (~ p105(X42) & r1(X11, X42)) | ? [X43] : (~ p205(X43) & r1(X11, X43))) & (? [X44] : (~ p105(X44) & r1(X11, X44)) | ? [X45] : (~ p305(X45) & r1(X11, X45))) & (? [X46] : (~ p105(X46) & r1(X11, X46)) | ? [X47] : (~ p405(X47) & r1(X11, X47))) & (? [X48] : (~ p105(X48) & r1(X11, X48)) | ~ p505(X11)) & (? [X49] : (~ p105(X49) & r1(X11, X49)) | ~ p605(X11)) & (? [X50] : (~ p205(X50) & r1(X11, X50)) | ? [X51] : (~ p305(X51) & r1(X11, X51))) & (? [X52] : (~ p205(X52) & r1(X11, X52)) | ? [X53] : (~ p405(X53) & r1(X11, X53))) & (? [X54] : (~ p205(X54) & r1(X11, X54)) | ~ p505(X11)) & (? [X55] : (~ p205(X55) & r1(X11, X55)) | ~ p605(X11)) & (? [X56] : (~ p305(X56) & r1(X11, X56)) | ? [X57] : (~ p405(X57) & r1(X11, X57))) & (? [X58] : (~ p305(X58) & r1(X11, X58)) | ~ p505(X11)) & (? [X59] : (~ p305(X59) & r1(X11, X59)) | ~ p605(X11)) & (? [X60] : (~ p405(X60) & r1(X11, X60)) | ~ p505(X11)) & (? [X61] : (~ p405(X61) & r1(X11, X61)) | ~ p605(X11)) & (~ p505(X11) | ~ p605(X11))) | ~ r1(X0, X11))), inference(flattening, [], [f6])).
fof(f6, plain, ? [X0] : (((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0))) & ! [X11] : (((~ p101(X11) | ~ p201(X11)) & (~ p101(X11) | ~ p301(X11)) & (~ p101(X11) | ~ p401(X11)) & (~ p101(X11) | ~ p501(X11)) & (~ p101(X11) | ~ p601(X11)) & (~ p201(X11) | ~ p301(X11)) & (~ p201(X11) | ~ p401(X11)) & (~ p201(X11) | ~ p501(X11)) & (~ p201(X11) | ~ p601(X11)) & (~ p301(X11) | ~ p401(X11)) & (~ p301(X11) | ~ p501(X11)) & (~ p301(X11) | ~ p601(X11)) & (~ p401(X11) | ~ p501(X11)) & (~ p401(X11) | ~ p601(X11)) & (~ p501(X11) | ~ p601(X11)) & (? [X12] : (~ p102(X12) & r1(X11, X12)) | ~ p202(X11)) & (? [X13] : (~ p102(X13) & r1(X11, X13)) | ~ p302(X11)) & (? [X14] : (~ p102(X14) & r1(X11, X14)) | ~ p402(X11)) & (? [X15] : (~ p102(X15) & r1(X11, X15)) | ~ p502(X11)) & (? [X16] : (~ p102(X16) & r1(X11, X16)) | ~ p602(X11)) & (~ p202(X11) | ~ p302(X11)) & (~ p202(X11) | ~ p402(X11)) & (~ p202(X11) | ~ p502(X11)) & (~ p202(X11) | ~ p602(X11)) & (~ p302(X11) | ~ p402(X11)) & (~ p302(X11) | ~ p502(X11)) & (~ p302(X11) | ~ p602(X11)) & (~ p402(X11) | ~ p502(X11)) & (~ p402(X11) | ~ p602(X11)) & (~ p502(X11) | ~ p602(X11)) & (? [X17] : (~ p103(X17) & r1(X11, X17)) | ? [X18] : (~ p203(X18) & r1(X11, X18))) & (? [X19] : (~ p103(X19) & r1(X11, X19)) | ~ p303(X11)) & (? [X20] : (~ p103(X20) & r1(X11, X20)) | ~ p403(X11)) & (? [X21] : (~ p103(X21) & r1(X11, X21)) | ~ p503(X11)) & (? [X22] : (~ p103(X22) & r1(X11, X22)) | ~ p603(X11)) & (? [X23] : (~ p203(X23) & r1(X11, X23)) | ~ p303(X11)) & (? [X24] : (~ p203(X24) & r1(X11, X24)) | ~ p403(X11)) & (? [X25] : (~ p203(X25) & r1(X11, X25)) | ~ p503(X11)) & (? [X26] : (~ p203(X26) & r1(X11, X26)) | ~ p603(X11)) & (~ p303(X11) | ~ p403(X11)) & (~ p303(X11) | ~ p503(X11)) & (~ p303(X11) | ~ p603(X11)) & (~ p403(X11) | ~ p503(X11)) & (~ p403(X11) | ~ p603(X11)) & (~ p503(X11) | ~ p603(X11)) & (? [X27] : (~ p104(X27) & r1(X11, X27)) | ? [X28] : (~ p204(X28) & r1(X11, X28))) & (? [X29] : (~ p104(X29) & r1(X11, X29)) | ? [X30] : (~ p304(X30) & r1(X11, X30))) & (? [X31] : (~ p104(X31) & r1(X11, X31)) | ~ p404(X11)) & (? [X32] : (~ p104(X32) & r1(X11, X32)) | ~ p504(X11)) & (? [X33] : (~ p104(X33) & r1(X11, X33)) | ~ p604(X11)) & (? [X34] : (~ p204(X34) & r1(X11, X34)) | ? [X35] : (~ p304(X35) & r1(X11, X35))) & (? [X36] : (~ p204(X36) & r1(X11, X36)) | ~ p404(X11)) & (? [X37] : (~ p204(X37) & r1(X11, X37)) | ~ p504(X11)) & (? [X38] : (~ p204(X38) & r1(X11, X38)) | ~ p604(X11)) & (? [X39] : (~ p304(X39) & r1(X11, X39)) | ~ p404(X11)) & (? [X40] : (~ p304(X40) & r1(X11, X40)) | ~ p504(X11)) & (? [X41] : (~ p304(X41) & r1(X11, X41)) | ~ p604(X11)) & (~ p404(X11) | ~ p504(X11)) & (~ p404(X11) | ~ p604(X11)) & (~ p504(X11) | ~ p604(X11)) & (? [X42] : (~ p105(X42) & r1(X11, X42)) | ? [X43] : (~ p205(X43) & r1(X11, X43))) & (? [X44] : (~ p105(X44) & r1(X11, X44)) | ? [X45] : (~ p305(X45) & r1(X11, X45))) & (? [X46] : (~ p105(X46) & r1(X11, X46)) | ? [X47] : (~ p405(X47) & r1(X11, X47))) & (? [X48] : (~ p105(X48) & r1(X11, X48)) | ~ p505(X11)) & (? [X49] : (~ p105(X49) & r1(X11, X49)) | ~ p605(X11)) & (? [X50] : (~ p205(X50) & r1(X11, X50)) | ? [X51] : (~ p305(X51) & r1(X11, X51))) & (? [X52] : (~ p205(X52) & r1(X11, X52)) | ? [X53] : (~ p405(X53) & r1(X11, X53))) & (? [X54] : (~ p205(X54) & r1(X11, X54)) | ~ p505(X11)) & (? [X55] : (~ p205(X55) & r1(X11, X55)) | ~ p605(X11)) & (? [X56] : (~ p305(X56) & r1(X11, X56)) | ? [X57] : (~ p405(X57) & r1(X11, X57))) & (? [X58] : (~ p305(X58) & r1(X11, X58)) | ~ p505(X11)) & (? [X59] : (~ p305(X59) & r1(X11, X59)) | ~ p605(X11)) & (? [X60] : (~ p405(X60) & r1(X11, X60)) | ~ p505(X11)) & (? [X61] : (~ p405(X61) & r1(X11, X61)) | ~ p605(X11)) & (~ p505(X11) | ~ p605(X11))) | ~ r1(X0, X11))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0))) | ~ ! [X11] : (~ ((p101(X11) & p201(X11)) | (p101(X11) & p301(X11)) | (p101(X11) & p401(X11)) | (p101(X11) & p501(X11)) | (p101(X11) & p601(X11)) | (p201(X11) & p301(X11)) | (p201(X11) & p401(X11)) | (p201(X11) & p501(X11)) | (p201(X11) & p601(X11)) | (p301(X11) & p401(X11)) | (p301(X11) & p501(X11)) | (p301(X11) & p601(X11)) | (p401(X11) & p501(X11)) | (p401(X11) & p601(X11)) | (p501(X11) & p601(X11)) | (! [X12] : (p102(X12) | ~ r1(X11, X12)) & p202(X11)) | (! [X13] : (p102(X13) | ~ r1(X11, X13)) & p302(X11)) | (! [X14] : (p102(X14) | ~ r1(X11, X14)) & p402(X11)) | (! [X15] : (p102(X15) | ~ r1(X11, X15)) & p502(X11)) | (! [X16] : (p102(X16) | ~ r1(X11, X16)) & p602(X11)) | (p202(X11) & p302(X11)) | (p202(X11) & p402(X11)) | (p202(X11) & p502(X11)) | (p202(X11) & p602(X11)) | (p302(X11) & p402(X11)) | (p302(X11) & p502(X11)) | (p302(X11) & p602(X11)) | (p402(X11) & p502(X11)) | (p402(X11) & p602(X11)) | (p502(X11) & p602(X11)) | (! [X17] : (p103(X17) | ~ r1(X11, X17)) & ! [X18] : (p203(X18) | ~ r1(X11, X18))) | (! [X19] : (p103(X19) | ~ r1(X11, X19)) & p303(X11)) | (! [X20] : (p103(X20) | ~ r1(X11, X20)) & p403(X11)) | (! [X21] : (p103(X21) | ~ r1(X11, X21)) & p503(X11)) | (! [X22] : (p103(X22) | ~ r1(X11, X22)) & p603(X11)) | (! [X23] : (p203(X23) | ~ r1(X11, X23)) & p303(X11)) | (! [X24] : (p203(X24) | ~ r1(X11, X24)) & p403(X11)) | (! [X25] : (p203(X25) | ~ r1(X11, X25)) & p503(X11)) | (! [X26] : (p203(X26) | ~ r1(X11, X26)) & p603(X11)) | (p303(X11) & p403(X11)) | (p303(X11) & p503(X11)) | (p303(X11) & p603(X11)) | (p403(X11) & p503(X11)) | (p403(X11) & p603(X11)) | (p503(X11) & p603(X11)) | (! [X27] : (p104(X27) | ~ r1(X11, X27)) & ! [X28] : (p204(X28) | ~ r1(X11, X28))) | (! [X29] : (p104(X29) | ~ r1(X11, X29)) & ! [X30] : (p304(X30) | ~ r1(X11, X30))) | (! [X31] : (p104(X31) | ~ r1(X11, X31)) & p404(X11)) | (! [X32] : (p104(X32) | ~ r1(X11, X32)) & p504(X11)) | (! [X33] : (p104(X33) | ~ r1(X11, X33)) & p604(X11)) | (! [X34] : (p204(X34) | ~ r1(X11, X34)) & ! [X35] : (p304(X35) | ~ r1(X11, X35))) | (! [X36] : (p204(X36) | ~ r1(X11, X36)) & p404(X11)) | (! [X37] : (p204(X37) | ~ r1(X11, X37)) & p504(X11)) | (! [X38] : (p204(X38) | ~ r1(X11, X38)) & p604(X11)) | (! [X39] : (p304(X39) | ~ r1(X11, X39)) & p404(X11)) | (! [X40] : (p304(X40) | ~ r1(X11, X40)) & p504(X11)) | (! [X41] : (p304(X41) | ~ r1(X11, X41)) & p604(X11)) | (p404(X11) & p504(X11)) | (p404(X11) & p604(X11)) | (p504(X11) & p604(X11)) | (! [X42] : (p105(X42) | ~ r1(X11, X42)) & ! [X43] : (p205(X43) | ~ r1(X11, X43))) | (! [X44] : (p105(X44) | ~ r1(X11, X44)) & ! [X45] : (p305(X45) | ~ r1(X11, X45))) | (! [X46] : (p105(X46) | ~ r1(X11, X46)) & ! [X47] : (p405(X47) | ~ r1(X11, X47))) | (! [X48] : (p105(X48) | ~ r1(X11, X48)) & p505(X11)) | (! [X49] : (p105(X49) | ~ r1(X11, X49)) & p605(X11)) | (! [X50] : (p205(X50) | ~ r1(X11, X50)) & ! [X51] : (p305(X51) | ~ r1(X11, X51))) | (! [X52] : (p205(X52) | ~ r1(X11, X52)) & ! [X53] : (p405(X53) | ~ r1(X11, X53))) | (! [X54] : (p205(X54) | ~ r1(X11, X54)) & p505(X11)) | (! [X55] : (p205(X55) | ~ r1(X11, X55)) & p605(X11)) | (! [X56] : (p305(X56) | ~ r1(X11, X56)) & ! [X57] : (p405(X57) | ~ r1(X11, X57))) | (! [X58] : (p305(X58) | ~ r1(X11, X58)) & p505(X11)) | (! [X59] : (p305(X59) | ~ r1(X11, X59)) & p605(X11)) | (! [X60] : (p405(X60) | ~ r1(X11, X60)) & p505(X11)) | (! [X61] : (p405(X61) | ~ r1(X11, X61)) & p605(X11)) | (p505(X11) & p605(X11))) | ~ r1(X0, X11))), inference(flattening, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0))) | ~ ! [X11] : (~ ((p101(X11) & p201(X11)) | (p101(X11) & p301(X11)) | (p101(X11) & p401(X11)) | (p101(X11) & p501(X11)) | (p101(X11) & p601(X11)) | (p201(X11) & p301(X11)) | (p201(X11) & p401(X11)) | (p201(X11) & p501(X11)) | (p201(X11) & p601(X11)) | (p301(X11) & p401(X11)) | (p301(X11) & p501(X11)) | (p301(X11) & p601(X11)) | (p401(X11) & p501(X11)) | (p401(X11) & p601(X11)) | (p501(X11) & p601(X11)) | (! [X12] : (p102(X12) | ~ r1(X11, X12)) & p202(X11)) | (! [X13] : (p102(X13) | ~ r1(X11, X13)) & p302(X11)) | (! [X14] : (p102(X14) | ~ r1(X11, X14)) & p402(X11)) | (! [X15] : (p102(X15) | ~ r1(X11, X15)) & p502(X11)) | (! [X16] : (p102(X16) | ~ r1(X11, X16)) & p602(X11)) | (p202(X11) & p302(X11)) | (p202(X11) & p402(X11)) | (p202(X11) & p502(X11)) | (p202(X11) & p602(X11)) | (p302(X11) & p402(X11)) | (p302(X11) & p502(X11)) | (p302(X11) & p602(X11)) | (p402(X11) & p502(X11)) | (p402(X11) & p602(X11)) | (p502(X11) & p602(X11)) | (! [X17] : (p103(X17) | ~ r1(X11, X17)) & ! [X18] : (p203(X18) | ~ r1(X11, X18))) | (! [X19] : (p103(X19) | ~ r1(X11, X19)) & p303(X11)) | (! [X20] : (p103(X20) | ~ r1(X11, X20)) & p403(X11)) | (! [X21] : (p103(X21) | ~ r1(X11, X21)) & p503(X11)) | (! [X22] : (p103(X22) | ~ r1(X11, X22)) & p603(X11)) | (! [X23] : (p203(X23) | ~ r1(X11, X23)) & p303(X11)) | (! [X24] : (p203(X24) | ~ r1(X11, X24)) & p403(X11)) | (! [X25] : (p203(X25) | ~ r1(X11, X25)) & p503(X11)) | (! [X26] : (p203(X26) | ~ r1(X11, X26)) & p603(X11)) | (p303(X11) & p403(X11)) | (p303(X11) & p503(X11)) | (p303(X11) & p603(X11)) | (p403(X11) & p503(X11)) | (p403(X11) & p603(X11)) | (p503(X11) & p603(X11)) | (! [X27] : (p104(X27) | ~ r1(X11, X27)) & ! [X28] : (p204(X28) | ~ r1(X11, X28))) | (! [X29] : (p104(X29) | ~ r1(X11, X29)) & ! [X30] : (p304(X30) | ~ r1(X11, X30))) | (! [X31] : (p104(X31) | ~ r1(X11, X31)) & p404(X11)) | (! [X32] : (p104(X32) | ~ r1(X11, X32)) & p504(X11)) | (! [X33] : (p104(X33) | ~ r1(X11, X33)) & p604(X11)) | (! [X34] : (p204(X34) | ~ r1(X11, X34)) & ! [X35] : (p304(X35) | ~ r1(X11, X35))) | (! [X36] : (p204(X36) | ~ r1(X11, X36)) & p404(X11)) | (! [X37] : (p204(X37) | ~ r1(X11, X37)) & p504(X11)) | (! [X38] : (p204(X38) | ~ r1(X11, X38)) & p604(X11)) | (! [X39] : (p304(X39) | ~ r1(X11, X39)) & p404(X11)) | (! [X40] : (p304(X40) | ~ r1(X11, X40)) & p504(X11)) | (! [X41] : (p304(X41) | ~ r1(X11, X41)) & p604(X11)) | (p404(X11) & p504(X11)) | (p404(X11) & p604(X11)) | (p504(X11) & p604(X11)) | (! [X42] : (p105(X42) | ~ r1(X11, X42)) & ! [X43] : (p205(X43) | ~ r1(X11, X43))) | (! [X44] : (p105(X44) | ~ r1(X11, X44)) & ! [X45] : (p305(X45) | ~ r1(X11, X45))) | (! [X46] : (p105(X46) | ~ r1(X11, X46)) & ! [X47] : (p405(X47) | ~ r1(X11, X47))) | (! [X48] : (p105(X48) | ~ r1(X11, X48)) & p505(X11)) | (! [X49] : (p105(X49) | ~ r1(X11, X49)) & p605(X11)) | (! [X50] : (p205(X50) | ~ r1(X11, X50)) & ! [X51] : (p305(X51) | ~ r1(X11, X51))) | (! [X52] : (p205(X52) | ~ r1(X11, X52)) & ! [X53] : (p405(X53) | ~ r1(X11, X53))) | (! [X54] : (p205(X54) | ~ r1(X11, X54)) & p505(X11)) | (! [X55] : (p205(X55) | ~ r1(X11, X55)) & p605(X11)) | (! [X56] : (p305(X56) | ~ r1(X11, X56)) & ! [X57] : (p405(X57) | ~ r1(X11, X57))) | (! [X58] : (p305(X58) | ~ r1(X11, X58)) & p505(X11)) | (! [X59] : (p305(X59) | ~ r1(X11, X59)) & p605(X11)) | (! [X60] : (p405(X60) | ~ r1(X11, X60)) & p505(X11)) | (! [X61] : (p405(X61) | ~ r1(X11, X61)) & p605(X11)) | (p505(X11) & p605(X11))) | ~ r1(X0, X11))), inference(rectify, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X1] : (p103(X1) | ~ r1(X0, X1)) | ! [X1] : (p104(X1) | ~ r1(X0, X1)) | ! [X1] : (p105(X1) | ~ r1(X0, X1))) & (p201(X0) | p202(X0) | ! [X1] : (p203(X1) | ~ r1(X0, X1)) | ! [X1] : (p204(X1) | ~ r1(X0, X1)) | ! [X1] : (p205(X1) | ~ r1(X0, X1))) & (p301(X0) | p302(X0) | p303(X0) | ! [X1] : (p304(X1) | ~ r1(X0, X1)) | ! [X1] : (p305(X1) | ~ r1(X0, X1))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X1] : (p405(X1) | ~ r1(X0, X1))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0))) | ~ ! [X1] : (~ ((p101(X1) & p201(X1)) | (p101(X1) & p301(X1)) | (p101(X1) & p401(X1)) | (p101(X1) & p501(X1)) | (p101(X1) & p601(X1)) | (p201(X1) & p301(X1)) | (p201(X1) & p401(X1)) | (p201(X1) & p501(X1)) | (p201(X1) & p601(X1)) | (p301(X1) & p401(X1)) | (p301(X1) & p501(X1)) | (p301(X1) & p601(X1)) | (p401(X1) & p501(X1)) | (p401(X1) & p601(X1)) | (p501(X1) & p601(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p202(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p302(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p402(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p502(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p602(X1)) | (p202(X1) & p302(X1)) | (p202(X1) & p402(X1)) | (p202(X1) & p502(X1)) | (p202(X1) & p602(X1)) | (p302(X1) & p402(X1)) | (p302(X1) & p502(X1)) | (p302(X1) & p602(X1)) | (p402(X1) & p502(X1)) | (p402(X1) & p602(X1)) | (p502(X1) & p602(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & ! [X0] : (p203(X0) | ~ r1(X1, X0))) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p303(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p403(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p503(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p603(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p303(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p403(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p503(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p603(X1)) | (p303(X1) & p403(X1)) | (p303(X1) & p503(X1)) | (p303(X1) & p603(X1)) | (p403(X1) & p503(X1)) | (p403(X1) & p603(X1)) | (p503(X1) & p603(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & ! [X0] : (p204(X0) | ~ r1(X1, X0))) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & ! [X0] : (p304(X0) | ~ r1(X1, X0))) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p604(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & ! [X0] : (p304(X0) | ~ r1(X1, X0))) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p604(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p604(X1)) | (p404(X1) & p504(X1)) | (p404(X1) & p604(X1)) | (p504(X1) & p604(X1)) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p205(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p305(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & ! [X0] : (p305(X0) | ~ r1(X1, X0))) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p405(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p405(X0) | ~ r1(X1, X0)) & p605(X1)) | (p505(X1) & p605(X1))) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X1] : (p103(X1) | ~ r1(X0, X1)) | ! [X1] : (p104(X1) | ~ r1(X0, X1)) | ! [X1] : (p105(X1) | ~ r1(X0, X1))) & (p201(X0) | p202(X0) | ! [X1] : (p203(X1) | ~ r1(X0, X1)) | ! [X1] : (p204(X1) | ~ r1(X0, X1)) | ! [X1] : (p205(X1) | ~ r1(X0, X1))) & (p301(X0) | p302(X0) | p303(X0) | ! [X1] : (p304(X1) | ~ r1(X0, X1)) | ! [X1] : (p305(X1) | ~ r1(X0, X1))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X1] : (p405(X1) | ~ r1(X0, X1))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0))) | ~ ! [X1] : (~ ((p101(X1) & p201(X1)) | (p101(X1) & p301(X1)) | (p101(X1) & p401(X1)) | (p101(X1) & p501(X1)) | (p101(X1) & p601(X1)) | (p201(X1) & p301(X1)) | (p201(X1) & p401(X1)) | (p201(X1) & p501(X1)) | (p201(X1) & p601(X1)) | (p301(X1) & p401(X1)) | (p301(X1) & p501(X1)) | (p301(X1) & p601(X1)) | (p401(X1) & p501(X1)) | (p401(X1) & p601(X1)) | (p501(X1) & p601(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p202(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p302(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p402(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p502(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p602(X1)) | (p202(X1) & p302(X1)) | (p202(X1) & p402(X1)) | (p202(X1) & p502(X1)) | (p202(X1) & p602(X1)) | (p302(X1) & p402(X1)) | (p302(X1) & p502(X1)) | (p302(X1) & p602(X1)) | (p402(X1) & p502(X1)) | (p402(X1) & p602(X1)) | (p502(X1) & p602(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & ! [X0] : (p203(X0) | ~ r1(X1, X0))) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p303(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p403(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p503(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p603(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p303(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p403(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p503(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p603(X1)) | (p303(X1) & p403(X1)) | (p303(X1) & p503(X1)) | (p303(X1) & p603(X1)) | (p403(X1) & p503(X1)) | (p403(X1) & p603(X1)) | (p503(X1) & p603(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & ! [X0] : (p204(X0) | ~ r1(X1, X0))) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & ! [X0] : (p304(X0) | ~ r1(X1, X0))) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p604(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & ! [X0] : (p304(X0) | ~ r1(X1, X0))) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p604(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p604(X1)) | (p404(X1) & p504(X1)) | (p404(X1) & p604(X1)) | (p504(X1) & p604(X1)) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p205(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p305(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & ! [X0] : (p305(X0) | ~ r1(X1, X0))) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p405(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p405(X0) | ~ r1(X1, X0)) & p605(X1)) | (p505(X1) & p605(X1))) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL666+1.005.p', main)).
fof(f16560, plain, (~ sP13(sK91) | (~ spl92_21 | ~ spl92_1411)), inference(subsumption_resolution, [], [f16559, f484])).
fof(f484, plain, (p505(sK91) | ~ spl92_21), inference(avatar_component_clause, [], [f482])).
fof(f482, plain, (spl92_21 <=> p505(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_21])])).
fof(f16559, plain, (~ p505(sK91) | ~ sP13(sK91) | ~ spl92_1411), inference(resolution, [], [f16509, f353])).
fof(f353, plain, ! [X0] : (~ p305(sK67(X0)) | ~ p505(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ! [X0] : ((~ p305(sK67(X0)) & r1(X0, sK67(X0))) | ~ p505(X0) | ~ sP13(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK67])], [f157, f158])).
fof(f158, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) => (~ p305(sK67(X0)) & r1(X0, sK67(X0)))), introduced(choice_axiom, [])).
fof(f157, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) | ~ p505(X0) | ~ sP13(X0)), inference(rectify, [], [f156])).
fof(f156, plain, ! [X11] : (? [X58] : (~ p305(X58) & r1(X11, X58)) | ~ p505(X11) | ~ sP13(X11)), inference(nnf_transformation, [], [f21])).
fof(f16509, plain, (p305(sK67(sK91)) | ~ spl92_1411), inference(avatar_component_clause, [], [f16507])).
fof(f16507, plain, (spl92_1411 <=> p305(sK67(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1411])])).
fof(f16510, plain, (~ spl92_21 | spl92_1411 | ~ spl92_11), inference(avatar_split_clause, [], [f16505, f443, f16507, f482])).
fof(f443, plain, (spl92_11 <=> ! [X9] : (p305(X9) | ~ r1(sK91, X9))), introduced(avatar_definition, [new_symbols(naming, [spl92_11])])).
fof(f16505, plain, (p305(sK67(sK91)) | ~ p505(sK91) | ~ spl92_11), inference(subsumption_resolution, [], [f16088, f527])).
fof(f16088, plain, (p305(sK67(sK91)) | ~ p505(sK91) | ~ sP13(sK91) | ~ spl92_11), inference(resolution, [], [f444, f352])).
fof(f352, plain, ! [X0] : (r1(X0, sK67(X0)) | ~ p505(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f159])).
fof(f444, plain, (! [X9] : (~ r1(sK91, X9) | p305(X9)) | ~ spl92_11), inference(avatar_component_clause, [], [f443])).
fof(f16472, plain, (spl92_581 | spl92_463 | ~ spl92_16), inference(avatar_split_clause, [], [f16471, f462, f4597, f5636])).
fof(f5636, plain, (spl92_581 <=> r1(sK91, sK89(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_581])])).
fof(f4597, plain, (spl92_463 <=> p405(sK90(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_463])])).
fof(f462, plain, (spl92_16 <=> ! [X10] : (p405(X10) | ~ r1(sK91, X10))), introduced(avatar_definition, [new_symbols(naming, [spl92_16])])).
fof(f16471, plain, (p405(sK90(sK91)) | r1(sK91, sK89(sK91)) | ~ spl92_16), inference(subsumption_resolution, [], [f16450, f528])).
fof(f528, plain, sP0(sK91), inference(resolution, [], [f523, f230])).
fof(f230, plain, ! [X0] : (~ sP40(X0) | sP0(X0)), inference(cnf_transformation, [], [f51])).
fof(f16450, plain, (p405(sK90(sK91)) | r1(sK91, sK89(sK91)) | ~ sP0(sK91) | ~ spl92_16), inference(resolution, [], [f463, f396])).
fof(f396, plain, ! [X0] : (r1(X0, sK90(X0)) | r1(X0, sK89(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f221])).
fof(f221, plain, ! [X0] : ((~ p305(sK89(X0)) & r1(X0, sK89(X0))) | (~ p405(sK90(X0)) & r1(X0, sK90(X0))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK89, sK90])], [f218, f220, f219])).
fof(f219, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) => (~ p305(sK89(X0)) & r1(X0, sK89(X0)))), introduced(choice_axiom, [])).
fof(f220, plain, ! [X0] : (? [X2] : (~ p405(X2) & r1(X0, X2)) => (~ p405(sK90(X0)) & r1(X0, sK90(X0)))), introduced(choice_axiom, [])).
fof(f218, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) | ? [X2] : (~ p405(X2) & r1(X0, X2)) | ~ sP0(X0)), inference(rectify, [], [f217])).
fof(f217, plain, ! [X11] : (? [X56] : (~ p305(X56) & r1(X11, X56)) | ? [X57] : (~ p405(X57) & r1(X11, X57)) | ~ sP0(X11)), inference(nnf_transformation, [], [f8])).
fof(f463, plain, (! [X10] : (~ r1(sK91, X10) | p405(X10)) | ~ spl92_16), inference(avatar_component_clause, [], [f462])).
fof(f16470, plain, (~ spl92_463 | spl92_462 | ~ spl92_11), inference(avatar_split_clause, [], [f16469, f443, f4593, f4597])).
fof(f4593, plain, (spl92_462 <=> p305(sK89(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_462])])).
fof(f16469, plain, (p305(sK89(sK91)) | ~ p405(sK90(sK91)) | ~ spl92_11), inference(subsumption_resolution, [], [f16119, f528])).
fof(f16119, plain, (p305(sK89(sK91)) | ~ p405(sK90(sK91)) | ~ sP0(sK91) | ~ spl92_11), inference(resolution, [], [f444, f397])).
fof(f397, plain, ! [X0] : (r1(X0, sK89(X0)) | ~ p405(sK90(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f221])).
fof(f16462, plain, (~ spl92_462 | spl92_463 | ~ spl92_16), inference(avatar_split_clause, [], [f16461, f462, f4597, f4593])).
fof(f16461, plain, (p405(sK90(sK91)) | ~ p305(sK89(sK91)) | ~ spl92_16), inference(subsumption_resolution, [], [f16451, f528])).
fof(f16451, plain, (p405(sK90(sK91)) | ~ p305(sK89(sK91)) | ~ sP0(sK91) | ~ spl92_16), inference(resolution, [], [f463, f398])).
fof(f398, plain, ! [X0] : (r1(X0, sK90(X0)) | ~ p305(sK89(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f221])).
fof(f16460, plain, (~ spl92_463 | ~ spl92_462), inference(avatar_split_clause, [], [f16459, f4593, f4597])).
fof(f16459, plain, (~ p405(sK90(sK91)) | ~ spl92_462), inference(subsumption_resolution, [], [f14520, f528])).
fof(f14520, plain, (~ p405(sK90(sK91)) | ~ sP0(sK91) | ~ spl92_462), inference(resolution, [], [f4595, f399])).
fof(f399, plain, ! [X0] : (~ p305(sK89(X0)) | ~ p405(sK90(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f221])).
fof(f4595, plain, (p305(sK89(sK91)) | ~ spl92_462), inference(avatar_component_clause, [], [f4593])).
fof(f16357, plain, (~ spl92_17 | ~ spl92_1394), inference(avatar_split_clause, [], [f16356, f16238, f465])).
fof(f465, plain, (spl92_17 <=> p404(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_17])])).
fof(f16238, plain, (spl92_1394 <=> p104(sK54(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1394])])).
fof(f16356, plain, (~ p404(sK91) | ~ spl92_1394), inference(subsumption_resolution, [], [f16352, f547])).
fof(f547, plain, sP26(sK91), inference(resolution, [], [f523, f252])).
fof(f252, plain, ! [X0] : (~ sP40(X0) | sP26(X0)), inference(cnf_transformation, [], [f51])).
fof(f16352, plain, (~ p404(sK91) | ~ sP26(sK91) | ~ spl92_1394), inference(resolution, [], [f16240, f327])).
fof(f327, plain, ! [X0] : (~ p104(sK54(X0)) | ~ p404(X0) | ~ sP26(X0)), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ! [X0] : ((~ p104(sK54(X0)) & r1(X0, sK54(X0))) | ~ p404(X0) | ~ sP26(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54])], [f105, f106])).
fof(f106, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) => (~ p104(sK54(X0)) & r1(X0, sK54(X0)))), introduced(choice_axiom, [])).
fof(f105, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) | ~ p404(X0) | ~ sP26(X0)), inference(rectify, [], [f104])).
fof(f104, plain, ! [X11] : (? [X31] : (~ p104(X31) & r1(X11, X31)) | ~ p404(X11) | ~ sP26(X11)), inference(nnf_transformation, [], [f34])).
fof(f16240, plain, (p104(sK54(sK91)) | ~ spl92_1394), inference(avatar_component_clause, [], [f16238])).
fof(f16267, plain, (spl92_516 | spl92_299 | ~ spl92_11), inference(avatar_split_clause, [], [f16266, f443, f3202, f5010])).
fof(f5010, plain, (spl92_516 <=> r1(sK91, sK81(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_516])])).
fof(f3202, plain, (spl92_299 <=> p305(sK82(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_299])])).
fof(f16266, plain, (p305(sK82(sK91)) | r1(sK91, sK81(sK91)) | ~ spl92_11), inference(subsumption_resolution, [], [f16108, f536])).
fof(f536, plain, sP4(sK91), inference(resolution, [], [f523, f238])).
fof(f238, plain, ! [X0] : (~ sP40(X0) | sP4(X0)), inference(cnf_transformation, [], [f51])).
fof(f16108, plain, (p305(sK82(sK91)) | r1(sK91, sK81(sK91)) | ~ sP4(sK91) | ~ spl92_11), inference(resolution, [], [f444, f380])).
fof(f380, plain, ! [X0] : (r1(X0, sK82(X0)) | r1(X0, sK81(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f201])).
fof(f201, plain, ! [X0] : ((~ p105(sK81(X0)) & r1(X0, sK81(X0))) | (~ p305(sK82(X0)) & r1(X0, sK82(X0))) | ~ sP4(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK81, sK82])], [f198, f200, f199])).
fof(f199, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) => (~ p105(sK81(X0)) & r1(X0, sK81(X0)))), introduced(choice_axiom, [])).
fof(f200, plain, ! [X0] : (? [X2] : (~ p305(X2) & r1(X0, X2)) => (~ p305(sK82(X0)) & r1(X0, sK82(X0)))), introduced(choice_axiom, [])).
fof(f198, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) | ? [X2] : (~ p305(X2) & r1(X0, X2)) | ~ sP4(X0)), inference(rectify, [], [f197])).
fof(f197, plain, ! [X11] : (? [X44] : (~ p105(X44) & r1(X11, X44)) | ? [X45] : (~ p305(X45) & r1(X11, X45)) | ~ sP4(X11)), inference(nnf_transformation, [], [f12])).
fof(f16241, plain, (~ spl92_17 | spl92_1394 | ~ spl92_2), inference(avatar_split_clause, [], [f16236, f411, f16238, f465])).
fof(f411, plain, (spl92_2 <=> ! [X3] : (p104(X3) | ~ r1(sK91, X3))), introduced(avatar_definition, [new_symbols(naming, [spl92_2])])).
fof(f16236, plain, (p104(sK54(sK91)) | ~ p404(sK91) | ~ spl92_2), inference(subsumption_resolution, [], [f15898, f547])).
fof(f15898, plain, (p104(sK54(sK91)) | ~ p404(sK91) | ~ sP26(sK91) | ~ spl92_2), inference(resolution, [], [f412, f326])).
fof(f326, plain, ! [X0] : (r1(X0, sK54(X0)) | ~ p404(X0) | ~ sP26(X0)), inference(cnf_transformation, [], [f107])).
fof(f412, plain, (! [X3] : (~ r1(sK91, X3) | p104(X3)) | ~ spl92_2), inference(avatar_component_clause, [], [f411])).
fof(f16205, plain, (spl92_539 | ~ spl92_11 | spl92_426), inference(avatar_split_clause, [], [f16204, f4290, f443, f5230])).
fof(f5230, plain, (spl92_539 <=> r1(sK91, sK85(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_539])])).
fof(f4290, plain, (spl92_426 <=> p305(sK86(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_426])])).
fof(f16204, plain, (r1(sK91, sK85(sK91)) | (~ spl92_11 | spl92_426)), inference(subsumption_resolution, [], [f16203, f532])).
fof(f532, plain, sP2(sK91), inference(resolution, [], [f523, f234])).
fof(f234, plain, ! [X0] : (~ sP40(X0) | sP2(X0)), inference(cnf_transformation, [], [f51])).
fof(f16203, plain, (r1(sK91, sK85(sK91)) | ~ sP2(sK91) | (~ spl92_11 | spl92_426)), inference(subsumption_resolution, [], [f16114, f4292])).
fof(f4292, plain, (~ p305(sK86(sK91)) | spl92_426), inference(avatar_component_clause, [], [f4290])).
fof(f16114, plain, (p305(sK86(sK91)) | r1(sK91, sK85(sK91)) | ~ sP2(sK91) | ~ spl92_11), inference(resolution, [], [f444, f388])).
fof(f388, plain, ! [X0] : (r1(X0, sK86(X0)) | r1(X0, sK85(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f211])).
fof(f211, plain, ! [X0] : ((~ p205(sK85(X0)) & r1(X0, sK85(X0))) | (~ p305(sK86(X0)) & r1(X0, sK86(X0))) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK85, sK86])], [f208, f210, f209])).
fof(f209, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) => (~ p205(sK85(X0)) & r1(X0, sK85(X0)))), introduced(choice_axiom, [])).
fof(f210, plain, ! [X0] : (? [X2] : (~ p305(X2) & r1(X0, X2)) => (~ p305(sK86(X0)) & r1(X0, sK86(X0)))), introduced(choice_axiom, [])).
fof(f208, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) | ? [X2] : (~ p305(X2) & r1(X0, X2)) | ~ sP2(X0)), inference(rectify, [], [f207])).
fof(f207, plain, ! [X11] : (? [X50] : (~ p205(X50) & r1(X11, X50)) | ? [X51] : (~ p305(X51) & r1(X11, X51)) | ~ sP2(X11)), inference(nnf_transformation, [], [f10])).
fof(f16144, plain, (~ spl92_11 | ~ spl92_26), inference(avatar_contradiction_clause, [], [f16143])).
fof(f16143, plain, ($false | (~ spl92_11 | ~ spl92_26)), inference(subsumption_resolution, [], [f16142, f526])).
fof(f526, plain, sP12(sK91), inference(resolution, [], [f523, f228])).
fof(f228, plain, ! [X0] : (~ sP40(X0) | sP12(X0)), inference(cnf_transformation, [], [f51])).
fof(f16142, plain, (~ sP12(sK91) | (~ spl92_11 | ~ spl92_26)), inference(subsumption_resolution, [], [f16141, f505])).
fof(f505, plain, (p605(sK91) | ~ spl92_26), inference(avatar_component_clause, [], [f503])).
fof(f503, plain, (spl92_26 <=> p605(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_26])])).
fof(f16141, plain, (~ p605(sK91) | ~ sP12(sK91) | (~ spl92_11 | ~ spl92_26)), inference(resolution, [], [f16133, f355])).
fof(f355, plain, ! [X0] : (~ p305(sK68(X0)) | ~ p605(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f163])).
fof(f163, plain, ! [X0] : ((~ p305(sK68(X0)) & r1(X0, sK68(X0))) | ~ p605(X0) | ~ sP12(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK68])], [f161, f162])).
fof(f162, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) => (~ p305(sK68(X0)) & r1(X0, sK68(X0)))), introduced(choice_axiom, [])).
fof(f161, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) | ~ p605(X0) | ~ sP12(X0)), inference(rectify, [], [f160])).
fof(f160, plain, ! [X11] : (? [X59] : (~ p305(X59) & r1(X11, X59)) | ~ p605(X11) | ~ sP12(X11)), inference(nnf_transformation, [], [f20])).
fof(f16133, plain, (p305(sK68(sK91)) | (~ spl92_11 | ~ spl92_26)), inference(subsumption_resolution, [], [f16132, f526])).
fof(f16132, plain, (p305(sK68(sK91)) | ~ sP12(sK91) | (~ spl92_11 | ~ spl92_26)), inference(subsumption_resolution, [], [f16089, f505])).
fof(f16089, plain, (p305(sK68(sK91)) | ~ p605(sK91) | ~ sP12(sK91) | ~ spl92_11), inference(resolution, [], [f444, f354])).
fof(f354, plain, ! [X0] : (r1(X0, sK68(X0)) | ~ p605(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f163])).
fof(f16038, plain, (~ spl92_6 | ~ spl92_26), inference(avatar_contradiction_clause, [], [f16037])).
fof(f16037, plain, ($false | (~ spl92_6 | ~ spl92_26)), inference(subsumption_resolution, [], [f16036, f529])).
fof(f529, plain, sP14(sK91), inference(resolution, [], [f523, f231])).
fof(f231, plain, ! [X0] : (~ sP40(X0) | sP14(X0)), inference(cnf_transformation, [], [f51])).
fof(f16036, plain, (~ sP14(sK91) | (~ spl92_6 | ~ spl92_26)), inference(subsumption_resolution, [], [f16035, f505])).
fof(f16035, plain, (~ p605(sK91) | ~ sP14(sK91) | (~ spl92_6 | ~ spl92_26)), inference(resolution, [], [f16029, f351])).
fof(f351, plain, ! [X0] : (~ p205(sK66(X0)) | ~ p605(X0) | ~ sP14(X0)), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ! [X0] : ((~ p205(sK66(X0)) & r1(X0, sK66(X0))) | ~ p605(X0) | ~ sP14(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK66])], [f153, f154])).
fof(f154, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) => (~ p205(sK66(X0)) & r1(X0, sK66(X0)))), introduced(choice_axiom, [])).
fof(f153, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) | ~ p605(X0) | ~ sP14(X0)), inference(rectify, [], [f152])).
fof(f152, plain, ! [X11] : (? [X55] : (~ p205(X55) & r1(X11, X55)) | ~ p605(X11) | ~ sP14(X11)), inference(nnf_transformation, [], [f22])).
fof(f16029, plain, (p205(sK66(sK91)) | (~ spl92_6 | ~ spl92_26)), inference(subsumption_resolution, [], [f16028, f529])).
fof(f16028, plain, (p205(sK66(sK91)) | ~ sP14(sK91) | (~ spl92_6 | ~ spl92_26)), inference(subsumption_resolution, [], [f15989, f505])).
fof(f15989, plain, (p205(sK66(sK91)) | ~ p605(sK91) | ~ sP14(sK91) | ~ spl92_6), inference(resolution, [], [f426, f350])).
fof(f350, plain, ! [X0] : (r1(X0, sK66(X0)) | ~ p605(X0) | ~ sP14(X0)), inference(cnf_transformation, [], [f155])).
fof(f426, plain, (! [X7] : (~ r1(sK91, X7) | p205(X7)) | ~ spl92_6), inference(avatar_component_clause, [], [f425])).
fof(f425, plain, (spl92_6 <=> ! [X7] : (p205(X7) | ~ r1(sK91, X7))), introduced(avatar_definition, [new_symbols(naming, [spl92_6])])).
fof(f15879, plain, (~ spl92_17 | ~ spl92_27), inference(avatar_split_clause, [], [f15878, f507, f465])).
fof(f507, plain, (spl92_27 <=> p604(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_27])])).
fof(f15878, plain, (~ p404(sK91) | ~ spl92_27), inference(subsumption_resolution, [], [f576, f509])).
fof(f509, plain, (p604(sK91) | ~ spl92_27), inference(avatar_component_clause, [], [f507])).
fof(f576, plain, (~ p604(sK91) | ~ p404(sK91)), inference(resolution, [], [f241, f523])).
fof(f241, plain, ! [X0] : (~ sP40(X0) | ~ p604(X0) | ~ p404(X0)), inference(cnf_transformation, [], [f51])).
fof(f15877, plain, (~ spl92_10 | ~ spl92_15), inference(avatar_split_clause, [], [f616, f457, f438])).
fof(f438, plain, (spl92_10 <=> p201(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_10])])).
fof(f457, plain, (spl92_15 <=> p301(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_15])])).
fof(f616, plain, (~ p301(sK91) | ~ p201(sK91)), inference(resolution, [], [f294, f523])).
fof(f294, plain, ! [X0] : (~ sP40(X0) | ~ p301(X0) | ~ p201(X0)), inference(cnf_transformation, [], [f51])).
fof(f15856, plain, (~ spl92_8 | ~ spl92_13), inference(avatar_contradiction_clause, [], [f15855])).
fof(f15855, plain, ($false | (~ spl92_8 | ~ spl92_13)), inference(subsumption_resolution, [], [f15854, f553])).
fof(f553, plain, sP30(sK91), inference(resolution, [], [f523, f264])).
fof(f264, plain, ! [X0] : (~ sP40(X0) | sP30(X0)), inference(cnf_transformation, [], [f51])).
fof(f15854, plain, (~ sP30(sK91) | (~ spl92_8 | ~ spl92_13)), inference(subsumption_resolution, [], [f15853, f451])).
fof(f451, plain, (p303(sK91) | ~ spl92_13), inference(avatar_component_clause, [], [f449])).
fof(f449, plain, (spl92_13 <=> p303(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_13])])).
fof(f15853, plain, (~ p303(sK91) | ~ sP30(sK91) | (~ spl92_8 | ~ spl92_13)), inference(resolution, [], [f15831, f319])).
fof(f319, plain, ! [X0] : (~ p203(sK50(X0)) | ~ p303(X0) | ~ sP30(X0)), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ! [X0] : ((~ p203(sK50(X0)) & r1(X0, sK50(X0))) | ~ p303(X0) | ~ sP30(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK50])], [f89, f90])).
fof(f90, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) => (~ p203(sK50(X0)) & r1(X0, sK50(X0)))), introduced(choice_axiom, [])).
fof(f89, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) | ~ p303(X0) | ~ sP30(X0)), inference(rectify, [], [f88])).
fof(f88, plain, ! [X11] : (? [X23] : (~ p203(X23) & r1(X11, X23)) | ~ p303(X11) | ~ sP30(X11)), inference(nnf_transformation, [], [f38])).
fof(f15831, plain, (p203(sK50(sK91)) | (~ spl92_8 | ~ spl92_13)), inference(subsumption_resolution, [], [f15830, f553])).
fof(f15830, plain, (p203(sK50(sK91)) | ~ sP30(sK91) | (~ spl92_8 | ~ spl92_13)), inference(subsumption_resolution, [], [f15777, f451])).
fof(f15777, plain, (p203(sK50(sK91)) | ~ p303(sK91) | ~ sP30(sK91) | ~ spl92_8), inference(resolution, [], [f432, f318])).
fof(f318, plain, ! [X0] : (r1(X0, sK50(X0)) | ~ p303(X0) | ~ sP30(X0)), inference(cnf_transformation, [], [f91])).
fof(f432, plain, (! [X5] : (~ r1(sK91, X5) | p203(X5)) | ~ spl92_8), inference(avatar_component_clause, [], [f431])).
fof(f431, plain, (spl92_8 <=> ! [X5] : (p203(X5) | ~ r1(sK91, X5))), introduced(avatar_definition, [new_symbols(naming, [spl92_8])])).
fof(f15695, plain, (~ spl92_3 | ~ spl92_13), inference(avatar_contradiction_clause, [], [f15694])).
fof(f15694, plain, ($false | (~ spl92_3 | ~ spl92_13)), inference(subsumption_resolution, [], [f15693, f557])).
fof(f557, plain, sP34(sK91), inference(resolution, [], [f523, f268])).
fof(f268, plain, ! [X0] : (~ sP40(X0) | sP34(X0)), inference(cnf_transformation, [], [f51])).
fof(f15693, plain, (~ sP34(sK91) | (~ spl92_3 | ~ spl92_13)), inference(subsumption_resolution, [], [f15692, f451])).
fof(f15692, plain, (~ p303(sK91) | ~ sP34(sK91) | (~ spl92_3 | ~ spl92_13)), inference(resolution, [], [f15684, f311])).
fof(f311, plain, ! [X0] : (~ p103(sK46(X0)) | ~ p303(X0) | ~ sP34(X0)), inference(cnf_transformation, [], [f75])).
fof(f75, plain, ! [X0] : ((~ p103(sK46(X0)) & r1(X0, sK46(X0))) | ~ p303(X0) | ~ sP34(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK46])], [f73, f74])).
fof(f74, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) => (~ p103(sK46(X0)) & r1(X0, sK46(X0)))), introduced(choice_axiom, [])).
fof(f73, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) | ~ p303(X0) | ~ sP34(X0)), inference(rectify, [], [f72])).
fof(f72, plain, ! [X11] : (? [X19] : (~ p103(X19) & r1(X11, X19)) | ~ p303(X11) | ~ sP34(X11)), inference(nnf_transformation, [], [f42])).
fof(f15684, plain, (p103(sK46(sK91)) | (~ spl92_3 | ~ spl92_13)), inference(subsumption_resolution, [], [f15683, f557])).
fof(f15683, plain, (p103(sK46(sK91)) | ~ sP34(sK91) | (~ spl92_3 | ~ spl92_13)), inference(subsumption_resolution, [], [f15628, f451])).
fof(f15628, plain, (p103(sK46(sK91)) | ~ p303(sK91) | ~ sP34(sK91) | ~ spl92_3), inference(resolution, [], [f415, f310])).
fof(f310, plain, ! [X0] : (r1(X0, sK46(X0)) | ~ p303(X0) | ~ sP34(X0)), inference(cnf_transformation, [], [f75])).
fof(f415, plain, (! [X2] : (~ r1(sK91, X2) | p103(X2)) | ~ spl92_3), inference(avatar_component_clause, [], [f414])).
fof(f414, plain, (spl92_3 <=> ! [X2] : (p103(X2) | ~ r1(sK91, X2))), introduced(avatar_definition, [new_symbols(naming, [spl92_3])])).
fof(f15617, plain, (~ spl92_13 | ~ spl92_23), inference(avatar_split_clause, [], [f587, f490, f449])).
fof(f490, plain, (spl92_23 <=> p503(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_23])])).
fof(f587, plain, (~ p503(sK91) | ~ p303(sK91)), inference(resolution, [], [f259, f523])).
fof(f259, plain, ! [X0] : (~ sP40(X0) | ~ p503(X0) | ~ p303(X0)), inference(cnf_transformation, [], [f51])).
fof(f15616, plain, (~ spl92_4 | ~ spl92_19), inference(avatar_contradiction_clause, [], [f15615])).
fof(f15615, plain, ($false | (~ spl92_4 | ~ spl92_19)), inference(subsumption_resolution, [], [f15614, f561])).
fof(f561, plain, sP37(sK91), inference(resolution, [], [f523, f282])).
fof(f282, plain, ! [X0] : (~ sP40(X0) | sP37(X0)), inference(cnf_transformation, [], [f51])).
fof(f15614, plain, (~ sP37(sK91) | (~ spl92_4 | ~ spl92_19)), inference(subsumption_resolution, [], [f15613, f475])).
fof(f475, plain, (p402(sK91) | ~ spl92_19), inference(avatar_component_clause, [], [f473])).
fof(f473, plain, (spl92_19 <=> p402(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_19])])).
fof(f15613, plain, (~ p402(sK91) | ~ sP37(sK91) | (~ spl92_4 | ~ spl92_19)), inference(resolution, [], [f15592, f305])).
fof(f305, plain, ! [X0] : (~ p102(sK43(X0)) | ~ p402(X0) | ~ sP37(X0)), inference(cnf_transformation, [], [f63])).
fof(f63, plain, ! [X0] : ((~ p102(sK43(X0)) & r1(X0, sK43(X0))) | ~ p402(X0) | ~ sP37(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK43])], [f61, f62])).
fof(f62, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) => (~ p102(sK43(X0)) & r1(X0, sK43(X0)))), introduced(choice_axiom, [])).
fof(f61, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) | ~ p402(X0) | ~ sP37(X0)), inference(rectify, [], [f60])).
fof(f60, plain, ! [X11] : (? [X14] : (~ p102(X14) & r1(X11, X14)) | ~ p402(X11) | ~ sP37(X11)), inference(nnf_transformation, [], [f45])).
fof(f15592, plain, (p102(sK43(sK91)) | (~ spl92_4 | ~ spl92_19)), inference(subsumption_resolution, [], [f15591, f561])).
fof(f15591, plain, (p102(sK43(sK91)) | ~ sP37(sK91) | (~ spl92_4 | ~ spl92_19)), inference(subsumption_resolution, [], [f15533, f475])).
fof(f15533, plain, (p102(sK43(sK91)) | ~ p402(sK91) | ~ sP37(sK91) | ~ spl92_4), inference(resolution, [], [f418, f304])).
fof(f304, plain, ! [X0] : (r1(X0, sK43(X0)) | ~ p402(X0) | ~ sP37(X0)), inference(cnf_transformation, [], [f63])).
fof(f418, plain, (! [X1] : (~ r1(sK91, X1) | p102(X1)) | ~ spl92_4), inference(avatar_component_clause, [], [f417])).
fof(f417, plain, (spl92_4 <=> ! [X1] : (p102(X1) | ~ r1(sK91, X1))), introduced(avatar_definition, [new_symbols(naming, [spl92_4])])).
fof(f15524, plain, (~ spl92_27 | ~ spl92_371), inference(avatar_split_clause, [], [f15523, f3643, f507])).
fof(f3643, plain, (spl92_371 <=> p204(sK59(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_371])])).
fof(f15523, plain, (~ p604(sK91) | ~ spl92_371), inference(subsumption_resolution, [], [f15517, f541])).
fof(f541, plain, sP21(sK91), inference(resolution, [], [f523, f246])).
fof(f246, plain, ! [X0] : (~ sP40(X0) | sP21(X0)), inference(cnf_transformation, [], [f51])).
fof(f15517, plain, (~ p604(sK91) | ~ sP21(sK91) | ~ spl92_371), inference(resolution, [], [f3645, f337])).
fof(f337, plain, ! [X0] : (~ p204(sK59(X0)) | ~ p604(X0) | ~ sP21(X0)), inference(cnf_transformation, [], [f127])).
fof(f127, plain, ! [X0] : ((~ p204(sK59(X0)) & r1(X0, sK59(X0))) | ~ p604(X0) | ~ sP21(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK59])], [f125, f126])).
fof(f126, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) => (~ p204(sK59(X0)) & r1(X0, sK59(X0)))), introduced(choice_axiom, [])).
fof(f125, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) | ~ p604(X0) | ~ sP21(X0)), inference(rectify, [], [f124])).
fof(f124, plain, ! [X11] : (? [X38] : (~ p204(X38) & r1(X11, X38)) | ~ p604(X11) | ~ sP21(X11)), inference(nnf_transformation, [], [f29])).
fof(f3645, plain, (p204(sK59(sK91)) | ~ spl92_371), inference(avatar_component_clause, [], [f3643])).
fof(f15522, plain, (~ spl92_5 | ~ spl92_20), inference(avatar_split_clause, [], [f620, f477, f420])).
fof(f420, plain, (spl92_5 <=> p101(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_5])])).
fof(f477, plain, (spl92_20 <=> p401(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_20])])).
fof(f620, plain, (~ p401(sK91) | ~ p101(sK91)), inference(resolution, [], [f297, f523])).
fof(f297, plain, ! [X0] : (~ sP40(X0) | ~ p401(X0) | ~ p101(X0)), inference(cnf_transformation, [], [f51])).
fof(f15521, plain, (~ spl92_5 | ~ spl92_15), inference(avatar_split_clause, [], [f621, f457, f420])).
fof(f621, plain, (~ p301(sK91) | ~ p101(sK91)), inference(resolution, [], [f298, f523])).
fof(f298, plain, ! [X0] : (~ sP40(X0) | ~ p301(X0) | ~ p101(X0)), inference(cnf_transformation, [], [f51])).
fof(f15441, plain, (spl92_476 | ~ spl92_7 | spl92_255), inference(avatar_split_clause, [], [f15440, f2903, f428, f4685])).
fof(f4685, plain, (spl92_476 <=> r1(sK91, sK73(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_476])])).
fof(f428, plain, (spl92_7 <=> ! [X6] : (p204(X6) | ~ r1(sK91, X6))), introduced(avatar_definition, [new_symbols(naming, [spl92_7])])).
fof(f2903, plain, (spl92_255 <=> p204(sK74(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_255])])).
fof(f15440, plain, (r1(sK91, sK73(sK91)) | (~ spl92_7 | spl92_255)), inference(subsumption_resolution, [], [f15439, f549])).
fof(f549, plain, sP8(sK91), inference(resolution, [], [f523, f254])).
fof(f254, plain, ! [X0] : (~ sP40(X0) | sP8(X0)), inference(cnf_transformation, [], [f51])).
fof(f15439, plain, (r1(sK91, sK73(sK91)) | ~ sP8(sK91) | (~ spl92_7 | spl92_255)), inference(subsumption_resolution, [], [f14842, f2905])).
fof(f2905, plain, (~ p204(sK74(sK91)) | spl92_255), inference(avatar_component_clause, [], [f2903])).
fof(f14842, plain, (p204(sK74(sK91)) | r1(sK91, sK73(sK91)) | ~ sP8(sK91) | ~ spl92_7), inference(resolution, [], [f429, f364])).
fof(f364, plain, ! [X0] : (r1(X0, sK74(X0)) | r1(X0, sK73(X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ! [X0] : ((~ p104(sK73(X0)) & r1(X0, sK73(X0))) | (~ p204(sK74(X0)) & r1(X0, sK74(X0))) | ~ sP8(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK73, sK74])], [f178, f180, f179])).
fof(f179, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) => (~ p104(sK73(X0)) & r1(X0, sK73(X0)))), introduced(choice_axiom, [])).
fof(f180, plain, ! [X0] : (? [X2] : (~ p204(X2) & r1(X0, X2)) => (~ p204(sK74(X0)) & r1(X0, sK74(X0)))), introduced(choice_axiom, [])).
fof(f178, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) | ? [X2] : (~ p204(X2) & r1(X0, X2)) | ~ sP8(X0)), inference(rectify, [], [f177])).
fof(f177, plain, ! [X11] : (? [X27] : (~ p104(X27) & r1(X11, X27)) | ? [X28] : (~ p204(X28) & r1(X11, X28)) | ~ sP8(X11)), inference(nnf_transformation, [], [f16])).
fof(f429, plain, (! [X6] : (~ r1(sK91, X6) | p204(X6)) | ~ spl92_7), inference(avatar_component_clause, [], [f428])).
fof(f15436, plain, (~ spl92_27 | spl92_371 | ~ spl92_7), inference(avatar_split_clause, [], [f15435, f428, f3643, f507])).
fof(f15435, plain, (p204(sK59(sK91)) | ~ p604(sK91) | ~ spl92_7), inference(subsumption_resolution, [], [f14826, f541])).
fof(f14826, plain, (p204(sK59(sK91)) | ~ p604(sK91) | ~ sP21(sK91) | ~ spl92_7), inference(resolution, [], [f429, f336])).
fof(f336, plain, ! [X0] : (r1(X0, sK59(X0)) | ~ p604(X0) | ~ sP21(X0)), inference(cnf_transformation, [], [f127])).
fof(f15429, plain, (~ spl92_255 | ~ spl92_2 | spl92_261), inference(avatar_split_clause, [], [f15428, f2942, f411, f2903])).
fof(f2942, plain, (spl92_261 <=> p104(sK73(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_261])])).
fof(f15428, plain, (~ p204(sK74(sK91)) | (~ spl92_2 | spl92_261)), inference(subsumption_resolution, [], [f15422, f549])).
fof(f15422, plain, (~ p204(sK74(sK91)) | ~ sP8(sK91) | (~ spl92_2 | spl92_261)), inference(subsumption_resolution, [], [f15386, f2944])).
fof(f2944, plain, (~ p104(sK73(sK91)) | spl92_261), inference(avatar_component_clause, [], [f2942])).
fof(f15386, plain, (p104(sK73(sK91)) | ~ p204(sK74(sK91)) | ~ sP8(sK91) | ~ spl92_2), inference(resolution, [], [f412, f365])).
fof(f365, plain, ! [X0] : (r1(X0, sK73(X0)) | ~ p204(sK74(X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f181])).
fof(f15346, plain, (~ spl92_4 | ~ spl92_14), inference(avatar_contradiction_clause, [], [f15345])).
fof(f15345, plain, ($false | (~ spl92_4 | ~ spl92_14)), inference(subsumption_resolution, [], [f15344, f562])).
fof(f562, plain, sP38(sK91), inference(resolution, [], [f523, f283])).
fof(f283, plain, ! [X0] : (~ sP40(X0) | sP38(X0)), inference(cnf_transformation, [], [f51])).
fof(f15344, plain, (~ sP38(sK91) | (~ spl92_4 | ~ spl92_14)), inference(subsumption_resolution, [], [f15343, f455])).
fof(f455, plain, (p302(sK91) | ~ spl92_14), inference(avatar_component_clause, [], [f453])).
fof(f453, plain, (spl92_14 <=> p302(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_14])])).
fof(f15343, plain, (~ p302(sK91) | ~ sP38(sK91) | (~ spl92_4 | ~ spl92_14)), inference(resolution, [], [f15022, f303])).
fof(f303, plain, ! [X0] : (~ p102(sK42(X0)) | ~ p302(X0) | ~ sP38(X0)), inference(cnf_transformation, [], [f59])).
fof(f59, plain, ! [X0] : ((~ p102(sK42(X0)) & r1(X0, sK42(X0))) | ~ p302(X0) | ~ sP38(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42])], [f57, f58])).
fof(f58, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) => (~ p102(sK42(X0)) & r1(X0, sK42(X0)))), introduced(choice_axiom, [])).
fof(f57, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) | ~ p302(X0) | ~ sP38(X0)), inference(rectify, [], [f56])).
fof(f56, plain, ! [X11] : (? [X13] : (~ p102(X13) & r1(X11, X13)) | ~ p302(X11) | ~ sP38(X11)), inference(nnf_transformation, [], [f46])).
fof(f15022, plain, (p102(sK42(sK91)) | (~ spl92_4 | ~ spl92_14)), inference(subsumption_resolution, [], [f15021, f562])).
fof(f15021, plain, (p102(sK42(sK91)) | ~ sP38(sK91) | (~ spl92_4 | ~ spl92_14)), inference(subsumption_resolution, [], [f14962, f455])).
fof(f14962, plain, (p102(sK42(sK91)) | ~ p302(sK91) | ~ sP38(sK91) | ~ spl92_4), inference(resolution, [], [f418, f302])).
fof(f302, plain, ! [X0] : (r1(X0, sK42(X0)) | ~ p302(X0) | ~ sP38(X0)), inference(cnf_transformation, [], [f59])).
fof(f14931, plain, (~ spl92_17 | ~ spl92_324), inference(avatar_split_clause, [], [f14930, f3359, f465])).
fof(f3359, plain, (spl92_324 <=> p204(sK57(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_324])])).
fof(f14930, plain, (~ p404(sK91) | ~ spl92_324), inference(subsumption_resolution, [], [f14926, f543])).
fof(f543, plain, sP23(sK91), inference(resolution, [], [f523, f248])).
fof(f248, plain, ! [X0] : (~ sP40(X0) | sP23(X0)), inference(cnf_transformation, [], [f51])).
fof(f14926, plain, (~ p404(sK91) | ~ sP23(sK91) | ~ spl92_324), inference(resolution, [], [f3361, f333])).
fof(f333, plain, ! [X0] : (~ p204(sK57(X0)) | ~ p404(X0) | ~ sP23(X0)), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ! [X0] : ((~ p204(sK57(X0)) & r1(X0, sK57(X0))) | ~ p404(X0) | ~ sP23(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK57])], [f117, f118])).
fof(f118, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) => (~ p204(sK57(X0)) & r1(X0, sK57(X0)))), introduced(choice_axiom, [])).
fof(f117, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) | ~ p404(X0) | ~ sP23(X0)), inference(rectify, [], [f116])).
fof(f116, plain, ! [X11] : (? [X36] : (~ p204(X36) & r1(X11, X36)) | ~ p404(X11) | ~ sP23(X11)), inference(nnf_transformation, [], [f31])).
fof(f3361, plain, (p204(sK57(sK91)) | ~ spl92_324), inference(avatar_component_clause, [], [f3359])).
fof(f14923, plain, (~ spl92_17 | spl92_324 | ~ spl92_7), inference(avatar_split_clause, [], [f14922, f428, f3359, f465])).
fof(f14922, plain, (p204(sK57(sK91)) | ~ p404(sK91) | ~ spl92_7), inference(subsumption_resolution, [], [f14824, f543])).
fof(f14824, plain, (p204(sK57(sK91)) | ~ p404(sK91) | ~ sP23(sK91) | ~ spl92_7), inference(resolution, [], [f429, f332])).
fof(f332, plain, ! [X0] : (r1(X0, sK57(X0)) | ~ p404(X0) | ~ sP23(X0)), inference(cnf_transformation, [], [f119])).
fof(f14919, plain, (~ spl92_14 | ~ spl92_19), inference(avatar_split_clause, [], [f597, f473, f453])).
fof(f597, plain, (~ p402(sK91) | ~ p302(sK91)), inference(resolution, [], [f275, f523])).
fof(f275, plain, ! [X0] : (~ sP40(X0) | ~ p402(X0) | ~ p302(X0)), inference(cnf_transformation, [], [f51])).
fof(f14916, plain, (~ spl92_19 | ~ spl92_24), inference(avatar_split_clause, [], [f14915, f494, f473])).
fof(f494, plain, (spl92_24 <=> p502(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_24])])).
fof(f14915, plain, (~ p402(sK91) | ~ spl92_24), inference(subsumption_resolution, [], [f593, f496])).
fof(f496, plain, (p502(sK91) | ~ spl92_24), inference(avatar_component_clause, [], [f494])).
fof(f593, plain, (~ p502(sK91) | ~ p402(sK91)), inference(resolution, [], [f272, f523])).
fof(f272, plain, ! [X0] : (~ sP40(X0) | ~ p502(X0) | ~ p402(X0)), inference(cnf_transformation, [], [f51])).
fof(f14908, plain, (~ spl92_14 | ~ spl92_24), inference(avatar_split_clause, [], [f14907, f494, f453])).
fof(f14907, plain, (~ p302(sK91) | ~ spl92_24), inference(subsumption_resolution, [], [f596, f496])).
fof(f596, plain, (~ p502(sK91) | ~ p302(sK91)), inference(resolution, [], [f274, f523])).
fof(f274, plain, ! [X0] : (~ sP40(X0) | ~ p502(X0) | ~ p302(X0)), inference(cnf_transformation, [], [f51])).
fof(f14884, plain, (~ spl92_22 | ~ spl92_197), inference(avatar_split_clause, [], [f14883, f2363, f486])).
fof(f486, plain, (spl92_22 <=> p504(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_22])])).
fof(f2363, plain, (spl92_197 <=> p204(sK58(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_197])])).
fof(f14883, plain, (~ p504(sK91) | ~ spl92_197), inference(subsumption_resolution, [], [f14879, f542])).
fof(f542, plain, sP22(sK91), inference(resolution, [], [f523, f247])).
fof(f247, plain, ! [X0] : (~ sP40(X0) | sP22(X0)), inference(cnf_transformation, [], [f51])).
fof(f14879, plain, (~ p504(sK91) | ~ sP22(sK91) | ~ spl92_197), inference(resolution, [], [f2365, f335])).
fof(f335, plain, ! [X0] : (~ p204(sK58(X0)) | ~ p504(X0) | ~ sP22(X0)), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ! [X0] : ((~ p204(sK58(X0)) & r1(X0, sK58(X0))) | ~ p504(X0) | ~ sP22(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK58])], [f121, f122])).
fof(f122, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) => (~ p204(sK58(X0)) & r1(X0, sK58(X0)))), introduced(choice_axiom, [])).
fof(f121, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) | ~ p504(X0) | ~ sP22(X0)), inference(rectify, [], [f120])).
fof(f120, plain, ! [X11] : (? [X37] : (~ p204(X37) & r1(X11, X37)) | ~ p504(X11) | ~ sP22(X11)), inference(nnf_transformation, [], [f30])).
fof(f2365, plain, (p204(sK58(sK91)) | ~ spl92_197), inference(avatar_component_clause, [], [f2363])).
fof(f14782, plain, (~ spl92_255 | ~ spl92_261), inference(avatar_split_clause, [], [f14781, f2942, f2903])).
fof(f14781, plain, (~ p204(sK74(sK91)) | ~ spl92_261), inference(subsumption_resolution, [], [f12614, f549])).
fof(f12614, plain, (~ p204(sK74(sK91)) | ~ sP8(sK91) | ~ spl92_261), inference(resolution, [], [f2943, f367])).
fof(f367, plain, ! [X0] : (~ p104(sK73(X0)) | ~ p204(sK74(X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f181])).
fof(f2943, plain, (p104(sK73(sK91)) | ~ spl92_261), inference(avatar_component_clause, [], [f2942])).
fof(f14780, plain, (~ spl92_9 | ~ spl92_14), inference(avatar_split_clause, [], [f602, f453, f434])).
fof(f434, plain, (spl92_9 <=> p202(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_9])])).
fof(f602, plain, (~ p302(sK91) | ~ p202(sK91)), inference(resolution, [], [f279, f523])).
fof(f279, plain, ! [X0] : (~ sP40(X0) | ~ p302(X0) | ~ p202(X0)), inference(cnf_transformation, [], [f51])).
fof(f14773, plain, (~ spl92_13 | ~ spl92_18), inference(avatar_split_clause, [], [f588, f469, f449])).
fof(f469, plain, (spl92_18 <=> p403(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_18])])).
fof(f588, plain, (~ p403(sK91) | ~ p303(sK91)), inference(resolution, [], [f260, f523])).
fof(f260, plain, ! [X0] : (~ sP40(X0) | ~ p403(X0) | ~ p303(X0)), inference(cnf_transformation, [], [f51])).
fof(f14760, plain, (~ spl92_299 | ~ spl92_305), inference(avatar_split_clause, [], [f14759, f3241, f3202])).
fof(f3241, plain, (spl92_305 <=> p105(sK81(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_305])])).
fof(f14759, plain, (~ p305(sK82(sK91)) | ~ spl92_305), inference(subsumption_resolution, [], [f14734, f536])).
fof(f14734, plain, (~ p305(sK82(sK91)) | ~ sP4(sK91) | ~ spl92_305), inference(resolution, [], [f3242, f383])).
fof(f383, plain, ! [X0] : (~ p105(sK81(X0)) | ~ p305(sK82(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f201])).
fof(f3242, plain, (p105(sK81(sK91)) | ~ spl92_305), inference(avatar_component_clause, [], [f3241])).
fof(f14758, plain, (~ spl92_15 | ~ spl92_25), inference(avatar_split_clause, [], [f610, f498, f457])).
fof(f498, plain, (spl92_25 <=> p501(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_25])])).
fof(f610, plain, (~ p501(sK91) | ~ p301(sK91)), inference(resolution, [], [f289, f523])).
fof(f289, plain, ! [X0] : (~ sP40(X0) | ~ p501(X0) | ~ p301(X0)), inference(cnf_transformation, [], [f51])).
fof(f14757, plain, (~ spl92_17 | ~ spl92_22), inference(avatar_split_clause, [], [f577, f486, f465])).
fof(f577, plain, (~ p504(sK91) | ~ p404(sK91)), inference(resolution, [], [f242, f523])).
fof(f242, plain, ! [X0] : (~ sP40(X0) | ~ p504(X0) | ~ p404(X0)), inference(cnf_transformation, [], [f51])).
fof(f14738, plain, (~ spl92_15 | ~ spl92_20), inference(avatar_split_clause, [], [f611, f477, f457])).
fof(f611, plain, (~ p401(sK91) | ~ p301(sK91)), inference(resolution, [], [f290, f523])).
fof(f290, plain, ! [X0] : (~ sP40(X0) | ~ p401(X0) | ~ p301(X0)), inference(cnf_transformation, [], [f51])).
fof(f14733, plain, (~ spl92_299 | spl92_305 | ~ spl92_1), inference(avatar_split_clause, [], [f14732, f408, f3241, f3202])).
fof(f408, plain, (spl92_1 <=> ! [X4] : (p105(X4) | ~ r1(sK91, X4))), introduced(avatar_definition, [new_symbols(naming, [spl92_1])])).
fof(f14732, plain, (p105(sK81(sK91)) | ~ p305(sK82(sK91)) | ~ spl92_1), inference(subsumption_resolution, [], [f13332, f536])).
fof(f13332, plain, (p105(sK81(sK91)) | ~ p305(sK82(sK91)) | ~ sP4(sK91) | ~ spl92_1), inference(resolution, [], [f409, f381])).
fof(f381, plain, ! [X0] : (r1(X0, sK81(X0)) | ~ p305(sK82(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f201])).
fof(f409, plain, (! [X4] : (~ r1(sK91, X4) | p105(X4)) | ~ spl92_1), inference(avatar_component_clause, [], [f408])).
fof(f14723, plain, (~ spl92_305 | spl92_299 | ~ spl92_11), inference(avatar_split_clause, [], [f14722, f443, f3202, f3241])).
fof(f14722, plain, (p305(sK82(sK91)) | ~ p105(sK81(sK91)) | ~ spl92_11), inference(subsumption_resolution, [], [f14275, f536])).
fof(f14275, plain, (p305(sK82(sK91)) | ~ p105(sK81(sK91)) | ~ sP4(sK91) | ~ spl92_11), inference(resolution, [], [f444, f382])).
fof(f382, plain, ! [X0] : (r1(X0, sK82(X0)) | ~ p105(sK81(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f201])).
fof(f14716, plain, (spl92_305 | ~ spl92_1 | ~ spl92_516), inference(avatar_split_clause, [], [f14712, f5010, f408, f3241])).
fof(f14712, plain, (p105(sK81(sK91)) | (~ spl92_1 | ~ spl92_516)), inference(resolution, [], [f5012, f409])).
fof(f5012, plain, (r1(sK91, sK81(sK91)) | ~ spl92_516), inference(avatar_component_clause, [], [f5010])).
fof(f14315, plain, (spl92_462 | ~ spl92_11 | ~ spl92_581), inference(avatar_split_clause, [], [f14226, f5636, f443, f4593])).
fof(f14226, plain, (p305(sK89(sK91)) | (~ spl92_11 | ~ spl92_581)), inference(resolution, [], [f444, f5638])).
fof(f5638, plain, (r1(sK91, sK89(sK91)) | ~ spl92_581), inference(avatar_component_clause, [], [f5636])).
fof(f14213, plain, (~ spl92_426 | ~ spl92_427), inference(avatar_split_clause, [], [f14212, f4296, f4290])).
fof(f4296, plain, (spl92_427 <=> p205(sK85(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_427])])).
fof(f14212, plain, (~ p305(sK86(sK91)) | ~ spl92_427), inference(subsumption_resolution, [], [f11263, f532])).
fof(f11263, plain, (~ p305(sK86(sK91)) | ~ sP2(sK91) | ~ spl92_427), inference(resolution, [], [f4298, f391])).
fof(f391, plain, ! [X0] : (~ p205(sK85(X0)) | ~ p305(sK86(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f211])).
fof(f4298, plain, (p205(sK85(sK91)) | ~ spl92_427), inference(avatar_component_clause, [], [f4296])).
fof(f14173, plain, (~ spl92_9 | ~ spl92_19), inference(avatar_split_clause, [], [f601, f473, f434])).
fof(f601, plain, (~ p402(sK91) | ~ p202(sK91)), inference(resolution, [], [f278, f523])).
fof(f278, plain, ! [X0] : (~ sP40(X0) | ~ p402(X0) | ~ p202(X0)), inference(cnf_transformation, [], [f51])).
fof(f14164, plain, (~ spl92_8 | ~ spl92_18), inference(avatar_contradiction_clause, [], [f14163])).
fof(f14163, plain, ($false | (~ spl92_8 | ~ spl92_18)), inference(subsumption_resolution, [], [f14162, f552])).
fof(f552, plain, sP29(sK91), inference(resolution, [], [f523, f263])).
fof(f263, plain, ! [X0] : (~ sP40(X0) | sP29(X0)), inference(cnf_transformation, [], [f51])).
fof(f14162, plain, (~ sP29(sK91) | (~ spl92_8 | ~ spl92_18)), inference(subsumption_resolution, [], [f14161, f471])).
fof(f471, plain, (p403(sK91) | ~ spl92_18), inference(avatar_component_clause, [], [f469])).
fof(f14161, plain, (~ p403(sK91) | ~ sP29(sK91) | (~ spl92_8 | ~ spl92_18)), inference(resolution, [], [f13960, f321])).
fof(f321, plain, ! [X0] : (~ p203(sK51(X0)) | ~ p403(X0) | ~ sP29(X0)), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ! [X0] : ((~ p203(sK51(X0)) & r1(X0, sK51(X0))) | ~ p403(X0) | ~ sP29(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK51])], [f93, f94])).
fof(f94, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) => (~ p203(sK51(X0)) & r1(X0, sK51(X0)))), introduced(choice_axiom, [])).
fof(f93, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) | ~ p403(X0) | ~ sP29(X0)), inference(rectify, [], [f92])).
fof(f92, plain, ! [X11] : (? [X24] : (~ p203(X24) & r1(X11, X24)) | ~ p403(X11) | ~ sP29(X11)), inference(nnf_transformation, [], [f37])).
fof(f13960, plain, (p203(sK51(sK91)) | (~ spl92_8 | ~ spl92_18)), inference(subsumption_resolution, [], [f13959, f552])).
fof(f13959, plain, (p203(sK51(sK91)) | ~ sP29(sK91) | (~ spl92_8 | ~ spl92_18)), inference(subsumption_resolution, [], [f13907, f471])).
fof(f13907, plain, (p203(sK51(sK91)) | ~ p403(sK91) | ~ sP29(sK91) | ~ spl92_8), inference(resolution, [], [f432, f320])).
fof(f320, plain, ! [X0] : (r1(X0, sK51(X0)) | ~ p403(X0) | ~ sP29(X0)), inference(cnf_transformation, [], [f95])).
fof(f13888, plain, (~ spl92_287 | ~ spl92_1 | spl92_293), inference(avatar_split_clause, [], [f13887, f3163, f408, f3124])).
fof(f3124, plain, (spl92_287 <=> p205(sK80(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_287])])).
fof(f3163, plain, (spl92_293 <=> p105(sK79(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_293])])).
fof(f13887, plain, (~ p205(sK80(sK91)) | (~ spl92_1 | spl92_293)), inference(subsumption_resolution, [], [f13886, f537])).
fof(f537, plain, sP5(sK91), inference(resolution, [], [f523, f239])).
fof(f239, plain, ! [X0] : (~ sP40(X0) | sP5(X0)), inference(cnf_transformation, [], [f51])).
fof(f13886, plain, (~ p205(sK80(sK91)) | ~ sP5(sK91) | (~ spl92_1 | spl92_293)), inference(subsumption_resolution, [], [f13329, f3165])).
fof(f3165, plain, (~ p105(sK79(sK91)) | spl92_293), inference(avatar_component_clause, [], [f3163])).
fof(f13329, plain, (p105(sK79(sK91)) | ~ p205(sK80(sK91)) | ~ sP5(sK91) | ~ spl92_1), inference(resolution, [], [f409, f377])).
fof(f377, plain, ! [X0] : (r1(X0, sK79(X0)) | ~ p205(sK80(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f196])).
fof(f196, plain, ! [X0] : ((~ p105(sK79(X0)) & r1(X0, sK79(X0))) | (~ p205(sK80(X0)) & r1(X0, sK80(X0))) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK79, sK80])], [f193, f195, f194])).
fof(f194, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) => (~ p105(sK79(X0)) & r1(X0, sK79(X0)))), introduced(choice_axiom, [])).
fof(f195, plain, ! [X0] : (? [X2] : (~ p205(X2) & r1(X0, X2)) => (~ p205(sK80(X0)) & r1(X0, sK80(X0)))), introduced(choice_axiom, [])).
fof(f193, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) | ? [X2] : (~ p205(X2) & r1(X0, X2)) | ~ sP5(X0)), inference(rectify, [], [f192])).
fof(f192, plain, ! [X11] : (? [X42] : (~ p105(X42) & r1(X11, X42)) | ? [X43] : (~ p205(X43) & r1(X11, X43)) | ~ sP5(X11)), inference(nnf_transformation, [], [f13])).
fof(f13873, plain, (~ spl92_18 | ~ spl92_23), inference(avatar_split_clause, [], [f584, f490, f469])).
fof(f584, plain, (~ p503(sK91) | ~ p403(sK91)), inference(resolution, [], [f257, f523])).
fof(f257, plain, ! [X0] : (~ sP40(X0) | ~ p503(X0) | ~ p403(X0)), inference(cnf_transformation, [], [f51])).
fof(f13857, plain, (~ spl92_287 | ~ spl92_293), inference(avatar_contradiction_clause, [], [f13856])).
fof(f13856, plain, ($false | (~ spl92_287 | ~ spl92_293)), inference(subsumption_resolution, [], [f13855, f537])).
fof(f13855, plain, (~ sP5(sK91) | (~ spl92_287 | ~ spl92_293)), inference(subsumption_resolution, [], [f13854, f3125])).
fof(f3125, plain, (p205(sK80(sK91)) | ~ spl92_287), inference(avatar_component_clause, [], [f3124])).
fof(f13854, plain, (~ p205(sK80(sK91)) | ~ sP5(sK91) | ~ spl92_293), inference(resolution, [], [f3164, f379])).
fof(f379, plain, ! [X0] : (~ p105(sK79(X0)) | ~ p205(sK80(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f196])).
fof(f3164, plain, (p105(sK79(sK91)) | ~ spl92_293), inference(avatar_component_clause, [], [f3163])).
fof(f13657, plain, (spl92_287 | ~ spl92_6 | ~ spl92_293), inference(avatar_split_clause, [], [f13656, f3163, f425, f3124])).
fof(f13656, plain, (p205(sK80(sK91)) | (~ spl92_6 | ~ spl92_293)), inference(subsumption_resolution, [], [f13655, f537])).
fof(f13655, plain, (p205(sK80(sK91)) | ~ sP5(sK91) | (~ spl92_6 | ~ spl92_293)), inference(subsumption_resolution, [], [f13410, f3164])).
fof(f13410, plain, (p205(sK80(sK91)) | ~ p105(sK79(sK91)) | ~ sP5(sK91) | ~ spl92_6), inference(resolution, [], [f426, f378])).
fof(f378, plain, ! [X0] : (r1(X0, sK80(X0)) | ~ p105(sK79(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f196])).
fof(f13645, plain, (spl92_663 | ~ spl92_6 | spl92_287), inference(avatar_split_clause, [], [f13644, f3124, f425, f6499])).
fof(f6499, plain, (spl92_663 <=> r1(sK91, sK79(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_663])])).
fof(f13644, plain, (r1(sK91, sK79(sK91)) | (~ spl92_6 | spl92_287)), inference(subsumption_resolution, [], [f13643, f537])).
fof(f13643, plain, (r1(sK91, sK79(sK91)) | ~ sP5(sK91) | (~ spl92_6 | spl92_287)), inference(subsumption_resolution, [], [f13409, f3126])).
fof(f3126, plain, (~ p205(sK80(sK91)) | spl92_287), inference(avatar_component_clause, [], [f3124])).
fof(f13409, plain, (p205(sK80(sK91)) | r1(sK91, sK79(sK91)) | ~ sP5(sK91) | ~ spl92_6), inference(resolution, [], [f426, f376])).
fof(f376, plain, ! [X0] : (r1(X0, sK80(X0)) | r1(X0, sK79(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f196])).
fof(f13640, plain, (~ spl92_1 | spl92_293 | ~ spl92_663), inference(avatar_contradiction_clause, [], [f13639])).
fof(f13639, plain, ($false | (~ spl92_1 | spl92_293 | ~ spl92_663)), inference(subsumption_resolution, [], [f13636, f3165])).
fof(f13636, plain, (p105(sK79(sK91)) | (~ spl92_1 | ~ spl92_663)), inference(resolution, [], [f6501, f409])).
fof(f6501, plain, (r1(sK91, sK79(sK91)) | ~ spl92_663), inference(avatar_component_clause, [], [f6499])).
fof(f13277, plain, (~ spl92_23 | ~ spl92_1239), inference(avatar_contradiction_clause, [], [f13276])).
fof(f13276, plain, ($false | (~ spl92_23 | ~ spl92_1239)), inference(subsumption_resolution, [], [f13275, f551])).
fof(f551, plain, sP28(sK91), inference(resolution, [], [f523, f262])).
fof(f262, plain, ! [X0] : (~ sP40(X0) | sP28(X0)), inference(cnf_transformation, [], [f51])).
fof(f13275, plain, (~ sP28(sK91) | (~ spl92_23 | ~ spl92_1239)), inference(subsumption_resolution, [], [f13274, f492])).
fof(f492, plain, (p503(sK91) | ~ spl92_23), inference(avatar_component_clause, [], [f490])).
fof(f13274, plain, (~ p503(sK91) | ~ sP28(sK91) | ~ spl92_1239), inference(resolution, [], [f13248, f323])).
fof(f323, plain, ! [X0] : (~ p203(sK52(X0)) | ~ p503(X0) | ~ sP28(X0)), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ! [X0] : ((~ p203(sK52(X0)) & r1(X0, sK52(X0))) | ~ p503(X0) | ~ sP28(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK52])], [f97, f98])).
fof(f98, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) => (~ p203(sK52(X0)) & r1(X0, sK52(X0)))), introduced(choice_axiom, [])).
fof(f97, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) | ~ p503(X0) | ~ sP28(X0)), inference(rectify, [], [f96])).
fof(f96, plain, ! [X11] : (? [X25] : (~ p203(X25) & r1(X11, X25)) | ~ p503(X11) | ~ sP28(X11)), inference(nnf_transformation, [], [f36])).
fof(f13248, plain, (p203(sK52(sK91)) | ~ spl92_1239), inference(avatar_component_clause, [], [f13246])).
fof(f13246, plain, (spl92_1239 <=> p203(sK52(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1239])])).
fof(f13249, plain, (~ spl92_23 | spl92_1239 | ~ spl92_8), inference(avatar_split_clause, [], [f13244, f431, f13246, f490])).
fof(f13244, plain, (p203(sK52(sK91)) | ~ p503(sK91) | ~ spl92_8), inference(subsumption_resolution, [], [f13170, f551])).
fof(f13170, plain, (p203(sK52(sK91)) | ~ p503(sK91) | ~ sP28(sK91) | ~ spl92_8), inference(resolution, [], [f432, f322])).
fof(f322, plain, ! [X0] : (r1(X0, sK52(X0)) | ~ p503(X0) | ~ sP28(X0)), inference(cnf_transformation, [], [f99])).
fof(f13243, plain, (~ spl92_8 | ~ spl92_28), inference(avatar_contradiction_clause, [], [f13242])).
fof(f13242, plain, ($false | (~ spl92_8 | ~ spl92_28)), inference(subsumption_resolution, [], [f13241, f550])).
fof(f550, plain, sP27(sK91), inference(resolution, [], [f523, f261])).
fof(f261, plain, ! [X0] : (~ sP40(X0) | sP27(X0)), inference(cnf_transformation, [], [f51])).
fof(f13241, plain, (~ sP27(sK91) | (~ spl92_8 | ~ spl92_28)), inference(subsumption_resolution, [], [f13240, f513])).
fof(f513, plain, (p603(sK91) | ~ spl92_28), inference(avatar_component_clause, [], [f511])).
fof(f511, plain, (spl92_28 <=> p603(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_28])])).
fof(f13240, plain, (~ p603(sK91) | ~ sP27(sK91) | (~ spl92_8 | ~ spl92_28)), inference(resolution, [], [f13222, f325])).
fof(f325, plain, ! [X0] : (~ p203(sK53(X0)) | ~ p603(X0) | ~ sP27(X0)), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ! [X0] : ((~ p203(sK53(X0)) & r1(X0, sK53(X0))) | ~ p603(X0) | ~ sP27(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK53])], [f101, f102])).
fof(f102, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) => (~ p203(sK53(X0)) & r1(X0, sK53(X0)))), introduced(choice_axiom, [])).
fof(f101, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) | ~ p603(X0) | ~ sP27(X0)), inference(rectify, [], [f100])).
fof(f100, plain, ! [X11] : (? [X26] : (~ p203(X26) & r1(X11, X26)) | ~ p603(X11) | ~ sP27(X11)), inference(nnf_transformation, [], [f35])).
fof(f13222, plain, (p203(sK53(sK91)) | (~ spl92_8 | ~ spl92_28)), inference(subsumption_resolution, [], [f13221, f550])).
fof(f13221, plain, (p203(sK53(sK91)) | ~ sP27(sK91) | (~ spl92_8 | ~ spl92_28)), inference(subsumption_resolution, [], [f13171, f513])).
fof(f13171, plain, (p203(sK53(sK91)) | ~ p603(sK91) | ~ sP27(sK91) | ~ spl92_8), inference(resolution, [], [f432, f324])).
fof(f324, plain, ! [X0] : (r1(X0, sK53(X0)) | ~ p603(X0) | ~ sP27(X0)), inference(cnf_transformation, [], [f103])).
fof(f13080, plain, (~ spl92_27 | ~ spl92_368), inference(avatar_split_clause, [], [f13079, f3625, f507])).
fof(f3625, plain, (spl92_368 <=> p304(sK62(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_368])])).
fof(f13079, plain, (~ p604(sK91) | ~ spl92_368), inference(subsumption_resolution, [], [f13072, f538])).
fof(f538, plain, sP18(sK91), inference(resolution, [], [f523, f243])).
fof(f243, plain, ! [X0] : (~ sP40(X0) | sP18(X0)), inference(cnf_transformation, [], [f51])).
fof(f13072, plain, (~ p604(sK91) | ~ sP18(sK91) | ~ spl92_368), inference(resolution, [], [f3627, f343])).
fof(f343, plain, ! [X0] : (~ p304(sK62(X0)) | ~ p604(X0) | ~ sP18(X0)), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ! [X0] : ((~ p304(sK62(X0)) & r1(X0, sK62(X0))) | ~ p604(X0) | ~ sP18(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK62])], [f137, f138])).
fof(f138, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) => (~ p304(sK62(X0)) & r1(X0, sK62(X0)))), introduced(choice_axiom, [])).
fof(f137, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) | ~ p604(X0) | ~ sP18(X0)), inference(rectify, [], [f136])).
fof(f136, plain, ! [X11] : (? [X41] : (~ p304(X41) & r1(X11, X41)) | ~ p604(X11) | ~ sP18(X11)), inference(nnf_transformation, [], [f26])).
fof(f3627, plain, (p304(sK62(sK91)) | ~ spl92_368), inference(avatar_component_clause, [], [f3625])).
fof(f13078, plain, (~ spl92_10 | ~ spl92_25), inference(avatar_split_clause, [], [f614, f498, f438])).
fof(f614, plain, (~ p501(sK91) | ~ p201(sK91)), inference(resolution, [], [f292, f523])).
fof(f292, plain, ! [X0] : (~ sP40(X0) | ~ p501(X0) | ~ p201(X0)), inference(cnf_transformation, [], [f51])).
fof(f13077, plain, (~ spl92_5 | ~ spl92_25), inference(avatar_split_clause, [], [f619, f498, f420])).
fof(f619, plain, (~ p501(sK91) | ~ p101(sK91)), inference(resolution, [], [f296, f523])).
fof(f296, plain, ! [X0] : (~ sP40(X0) | ~ p501(X0) | ~ p101(X0)), inference(cnf_transformation, [], [f51])).
fof(f13076, plain, (~ spl92_5 | ~ spl92_10), inference(avatar_split_clause, [], [f622, f438, f420])).
fof(f622, plain, (~ p201(sK91) | ~ p101(sK91)), inference(resolution, [], [f299, f523])).
fof(f299, plain, ! [X0] : (~ sP40(X0) | ~ p201(X0) | ~ p101(X0)), inference(cnf_transformation, [], [f51])).
fof(f12999, plain, (~ spl92_4 | ~ spl92_24), inference(avatar_contradiction_clause, [], [f12998])).
fof(f12998, plain, ($false | (~ spl92_4 | ~ spl92_24)), inference(subsumption_resolution, [], [f12997, f560])).
fof(f560, plain, sP36(sK91), inference(resolution, [], [f523, f281])).
fof(f281, plain, ! [X0] : (~ sP40(X0) | sP36(X0)), inference(cnf_transformation, [], [f51])).
fof(f12997, plain, (~ sP36(sK91) | (~ spl92_4 | ~ spl92_24)), inference(subsumption_resolution, [], [f12996, f496])).
fof(f12996, plain, (~ p502(sK91) | ~ sP36(sK91) | (~ spl92_4 | ~ spl92_24)), inference(resolution, [], [f12981, f307])).
fof(f307, plain, ! [X0] : (~ p102(sK44(X0)) | ~ p502(X0) | ~ sP36(X0)), inference(cnf_transformation, [], [f67])).
fof(f67, plain, ! [X0] : ((~ p102(sK44(X0)) & r1(X0, sK44(X0))) | ~ p502(X0) | ~ sP36(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK44])], [f65, f66])).
fof(f66, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) => (~ p102(sK44(X0)) & r1(X0, sK44(X0)))), introduced(choice_axiom, [])).
fof(f65, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) | ~ p502(X0) | ~ sP36(X0)), inference(rectify, [], [f64])).
fof(f64, plain, ! [X11] : (? [X15] : (~ p102(X15) & r1(X11, X15)) | ~ p502(X11) | ~ sP36(X11)), inference(nnf_transformation, [], [f44])).
fof(f12981, plain, (p102(sK44(sK91)) | (~ spl92_4 | ~ spl92_24)), inference(subsumption_resolution, [], [f12980, f560])).
fof(f12980, plain, (p102(sK44(sK91)) | ~ sP36(sK91) | (~ spl92_4 | ~ spl92_24)), inference(subsumption_resolution, [], [f12923, f496])).
fof(f12923, plain, (p102(sK44(sK91)) | ~ p502(sK91) | ~ sP36(sK91) | ~ spl92_4), inference(resolution, [], [f418, f306])).
fof(f306, plain, ! [X0] : (r1(X0, sK44(X0)) | ~ p502(X0) | ~ sP36(X0)), inference(cnf_transformation, [], [f67])).
fof(f12904, plain, (~ spl92_4 | ~ spl92_29), inference(avatar_contradiction_clause, [], [f12903])).
fof(f12903, plain, ($false | (~ spl92_4 | ~ spl92_29)), inference(subsumption_resolution, [], [f12902, f559])).
fof(f559, plain, sP35(sK91), inference(resolution, [], [f523, f280])).
fof(f280, plain, ! [X0] : (~ sP40(X0) | sP35(X0)), inference(cnf_transformation, [], [f51])).
fof(f12902, plain, (~ sP35(sK91) | (~ spl92_4 | ~ spl92_29)), inference(subsumption_resolution, [], [f12901, f517])).
fof(f517, plain, (p602(sK91) | ~ spl92_29), inference(avatar_component_clause, [], [f515])).
fof(f515, plain, (spl92_29 <=> p602(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_29])])).
fof(f12901, plain, (~ p602(sK91) | ~ sP35(sK91) | (~ spl92_4 | ~ spl92_29)), inference(resolution, [], [f12877, f309])).
fof(f309, plain, ! [X0] : (~ p102(sK45(X0)) | ~ p602(X0) | ~ sP35(X0)), inference(cnf_transformation, [], [f71])).
fof(f71, plain, ! [X0] : ((~ p102(sK45(X0)) & r1(X0, sK45(X0))) | ~ p602(X0) | ~ sP35(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK45])], [f69, f70])).
fof(f70, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) => (~ p102(sK45(X0)) & r1(X0, sK45(X0)))), introduced(choice_axiom, [])).
fof(f69, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) | ~ p602(X0) | ~ sP35(X0)), inference(rectify, [], [f68])).
fof(f68, plain, ! [X11] : (? [X16] : (~ p102(X16) & r1(X11, X16)) | ~ p602(X11) | ~ sP35(X11)), inference(nnf_transformation, [], [f43])).
fof(f12877, plain, (p102(sK45(sK91)) | (~ spl92_4 | ~ spl92_29)), inference(subsumption_resolution, [], [f12876, f559])).
fof(f12876, plain, (p102(sK45(sK91)) | ~ sP35(sK91) | (~ spl92_4 | ~ spl92_29)), inference(subsumption_resolution, [], [f12820, f517])).
fof(f12820, plain, (p102(sK45(sK91)) | ~ p602(sK91) | ~ sP35(sK91) | ~ spl92_4), inference(resolution, [], [f418, f308])).
fof(f308, plain, ! [X0] : (r1(X0, sK45(X0)) | ~ p602(X0) | ~ sP35(X0)), inference(cnf_transformation, [], [f71])).
fof(f12810, plain, (~ spl92_401 | spl92_376 | ~ spl92_16), inference(avatar_split_clause, [], [f12809, f462, f3846, f4068])).
fof(f4068, plain, (spl92_401 <=> p105(sK83(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_401])])).
fof(f3846, plain, (spl92_376 <=> p405(sK84(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_376])])).
fof(f12809, plain, (p405(sK84(sK91)) | ~ p105(sK83(sK91)) | ~ spl92_16), inference(subsumption_resolution, [], [f10188, f535])).
fof(f535, plain, sP3(sK91), inference(resolution, [], [f523, f237])).
fof(f237, plain, ! [X0] : (~ sP40(X0) | sP3(X0)), inference(cnf_transformation, [], [f51])).
fof(f10188, plain, (p405(sK84(sK91)) | ~ p105(sK83(sK91)) | ~ sP3(sK91) | ~ spl92_16), inference(resolution, [], [f463, f386])).
fof(f386, plain, ! [X0] : (r1(X0, sK84(X0)) | ~ p105(sK83(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ! [X0] : ((~ p105(sK83(X0)) & r1(X0, sK83(X0))) | (~ p405(sK84(X0)) & r1(X0, sK84(X0))) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK83, sK84])], [f203, f205, f204])).
fof(f204, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) => (~ p105(sK83(X0)) & r1(X0, sK83(X0)))), introduced(choice_axiom, [])).
fof(f205, plain, ! [X0] : (? [X2] : (~ p405(X2) & r1(X0, X2)) => (~ p405(sK84(X0)) & r1(X0, sK84(X0)))), introduced(choice_axiom, [])).
fof(f203, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) | ? [X2] : (~ p405(X2) & r1(X0, X2)) | ~ sP3(X0)), inference(rectify, [], [f202])).
fof(f202, plain, ! [X11] : (? [X46] : (~ p105(X46) & r1(X11, X46)) | ? [X47] : (~ p405(X47) & r1(X11, X47)) | ~ sP3(X11)), inference(nnf_transformation, [], [f11])).
fof(f12808, plain, (~ spl92_266 | ~ spl92_272), inference(avatar_contradiction_clause, [], [f12807])).
fof(f12807, plain, ($false | (~ spl92_266 | ~ spl92_272)), inference(subsumption_resolution, [], [f12806, f548])).
fof(f548, plain, sP7(sK91), inference(resolution, [], [f523, f253])).
fof(f253, plain, ! [X0] : (~ sP40(X0) | sP7(X0)), inference(cnf_transformation, [], [f51])).
fof(f12806, plain, (~ sP7(sK91) | (~ spl92_266 | ~ spl92_272)), inference(subsumption_resolution, [], [f12805, f2979])).
fof(f2979, plain, (p304(sK76(sK91)) | ~ spl92_266), inference(avatar_component_clause, [], [f2978])).
fof(f2978, plain, (spl92_266 <=> p304(sK76(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_266])])).
fof(f12805, plain, (~ p304(sK76(sK91)) | ~ sP7(sK91) | ~ spl92_272), inference(resolution, [], [f3018, f371])).
fof(f371, plain, ! [X0] : (~ p104(sK75(X0)) | ~ p304(sK76(X0)) | ~ sP7(X0)), inference(cnf_transformation, [], [f186])).
fof(f186, plain, ! [X0] : ((~ p104(sK75(X0)) & r1(X0, sK75(X0))) | (~ p304(sK76(X0)) & r1(X0, sK76(X0))) | ~ sP7(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK75, sK76])], [f183, f185, f184])).
fof(f184, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) => (~ p104(sK75(X0)) & r1(X0, sK75(X0)))), introduced(choice_axiom, [])).
fof(f185, plain, ! [X0] : (? [X2] : (~ p304(X2) & r1(X0, X2)) => (~ p304(sK76(X0)) & r1(X0, sK76(X0)))), introduced(choice_axiom, [])).
fof(f183, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) | ? [X2] : (~ p304(X2) & r1(X0, X2)) | ~ sP7(X0)), inference(rectify, [], [f182])).
fof(f182, plain, ! [X11] : (? [X29] : (~ p104(X29) & r1(X11, X29)) | ? [X30] : (~ p304(X30) & r1(X11, X30)) | ~ sP7(X11)), inference(nnf_transformation, [], [f15])).
fof(f3018, plain, (p104(sK75(sK91)) | ~ spl92_272), inference(avatar_component_clause, [], [f3017])).
fof(f3017, plain, (spl92_272 <=> p104(sK75(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_272])])).
fof(f12601, plain, (spl92_272 | ~ spl92_2 | ~ spl92_480), inference(avatar_split_clause, [], [f12522, f4713, f411, f3017])).
fof(f4713, plain, (spl92_480 <=> r1(sK91, sK75(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_480])])).
fof(f12522, plain, (p104(sK75(sK91)) | (~ spl92_2 | ~ spl92_480)), inference(resolution, [], [f412, f4715])).
fof(f4715, plain, (r1(sK91, sK75(sK91)) | ~ spl92_480), inference(avatar_component_clause, [], [f4713])).
fof(f12600, plain, (spl92_272 | ~ spl92_2 | ~ spl92_266), inference(avatar_split_clause, [], [f12599, f2978, f411, f3017])).
fof(f12599, plain, (p104(sK75(sK91)) | (~ spl92_2 | ~ spl92_266)), inference(subsumption_resolution, [], [f12598, f548])).
fof(f12598, plain, (p104(sK75(sK91)) | ~ sP7(sK91) | (~ spl92_2 | ~ spl92_266)), inference(subsumption_resolution, [], [f12563, f2979])).
fof(f12563, plain, (p104(sK75(sK91)) | ~ p304(sK76(sK91)) | ~ sP7(sK91) | ~ spl92_2), inference(resolution, [], [f412, f369])).
fof(f369, plain, ! [X0] : (r1(X0, sK75(X0)) | ~ p304(sK76(X0)) | ~ sP7(X0)), inference(cnf_transformation, [], [f186])).
fof(f12588, plain, (~ spl92_2 | spl92_261 | ~ spl92_476), inference(avatar_contradiction_clause, [], [f12587])).
fof(f12587, plain, ($false | (~ spl92_2 | spl92_261 | ~ spl92_476)), inference(subsumption_resolution, [], [f12521, f2944])).
fof(f12521, plain, (p104(sK73(sK91)) | (~ spl92_2 | ~ spl92_476)), inference(resolution, [], [f412, f4687])).
fof(f4687, plain, (r1(sK91, sK73(sK91)) | ~ spl92_476), inference(avatar_component_clause, [], [f4685])).
fof(f12520, plain, (~ spl92_376 | ~ spl92_401), inference(avatar_contradiction_clause, [], [f12519])).
fof(f12519, plain, ($false | (~ spl92_376 | ~ spl92_401)), inference(subsumption_resolution, [], [f12518, f535])).
fof(f12518, plain, (~ sP3(sK91) | (~ spl92_376 | ~ spl92_401)), inference(subsumption_resolution, [], [f12517, f3847])).
fof(f3847, plain, (p405(sK84(sK91)) | ~ spl92_376), inference(avatar_component_clause, [], [f3846])).
fof(f12517, plain, (~ p405(sK84(sK91)) | ~ sP3(sK91) | ~ spl92_401), inference(resolution, [], [f4069, f387])).
fof(f387, plain, ! [X0] : (~ p105(sK83(X0)) | ~ p405(sK84(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f206])).
fof(f4069, plain, (p105(sK83(sK91)) | ~ spl92_401), inference(avatar_component_clause, [], [f4068])).
fof(f12516, plain, (spl92_401 | ~ spl92_1 | ~ spl92_518), inference(avatar_split_clause, [], [f12424, f5027, f408, f4068])).
fof(f5027, plain, (spl92_518 <=> r1(sK91, sK83(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_518])])).
fof(f12424, plain, (p105(sK83(sK91)) | (~ spl92_1 | ~ spl92_518)), inference(resolution, [], [f409, f5029])).
fof(f5029, plain, (r1(sK91, sK83(sK91)) | ~ spl92_518), inference(avatar_component_clause, [], [f5027])).
fof(f12515, plain, (spl92_401 | ~ spl92_1 | ~ spl92_376), inference(avatar_split_clause, [], [f12514, f3846, f408, f4068])).
fof(f12514, plain, (p105(sK83(sK91)) | (~ spl92_1 | ~ spl92_376)), inference(subsumption_resolution, [], [f12513, f535])).
fof(f12513, plain, (p105(sK83(sK91)) | ~ sP3(sK91) | (~ spl92_1 | ~ spl92_376)), inference(subsumption_resolution, [], [f12475, f3847])).
fof(f12475, plain, (p105(sK83(sK91)) | ~ p405(sK84(sK91)) | ~ sP3(sK91) | ~ spl92_1), inference(resolution, [], [f409, f385])).
fof(f385, plain, ! [X0] : (r1(X0, sK83(X0)) | ~ p405(sK84(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f206])).
fof(f12412, plain, (~ spl92_21 | ~ spl92_1140), inference(avatar_split_clause, [], [f12411, f11815, f482])).
fof(f11815, plain, (spl92_1140 <=> p405(sK69(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1140])])).
fof(f12411, plain, (~ p505(sK91) | ~ spl92_1140), inference(subsumption_resolution, [], [f12407, f525])).
fof(f525, plain, sP11(sK91), inference(resolution, [], [f523, f227])).
fof(f227, plain, ! [X0] : (~ sP40(X0) | sP11(X0)), inference(cnf_transformation, [], [f51])).
fof(f12407, plain, (~ p505(sK91) | ~ sP11(sK91) | ~ spl92_1140), inference(resolution, [], [f11817, f357])).
fof(f357, plain, ! [X0] : (~ p405(sK69(X0)) | ~ p505(X0) | ~ sP11(X0)), inference(cnf_transformation, [], [f167])).
fof(f167, plain, ! [X0] : ((~ p405(sK69(X0)) & r1(X0, sK69(X0))) | ~ p505(X0) | ~ sP11(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK69])], [f165, f166])).
fof(f166, plain, ! [X0] : (? [X1] : (~ p405(X1) & r1(X0, X1)) => (~ p405(sK69(X0)) & r1(X0, sK69(X0)))), introduced(choice_axiom, [])).
fof(f165, plain, ! [X0] : (? [X1] : (~ p405(X1) & r1(X0, X1)) | ~ p505(X0) | ~ sP11(X0)), inference(rectify, [], [f164])).
fof(f164, plain, ! [X11] : (? [X60] : (~ p405(X60) & r1(X11, X60)) | ~ p505(X11) | ~ sP11(X11)), inference(nnf_transformation, [], [f19])).
fof(f11817, plain, (p405(sK69(sK91)) | ~ spl92_1140), inference(avatar_component_clause, [], [f11815])).
fof(f11818, plain, (~ spl92_21 | spl92_1140 | ~ spl92_16), inference(avatar_split_clause, [], [f11813, f462, f11815, f482])).
fof(f11813, plain, (p405(sK69(sK91)) | ~ p505(sK91) | ~ spl92_16), inference(subsumption_resolution, [], [f10166, f525])).
fof(f10166, plain, (p405(sK69(sK91)) | ~ p505(sK91) | ~ sP11(sK91) | ~ spl92_16), inference(resolution, [], [f463, f356])).
fof(f356, plain, ! [X0] : (r1(X0, sK69(X0)) | ~ p505(X0) | ~ sP11(X0)), inference(cnf_transformation, [], [f167])).
fof(f11812, plain, (~ spl92_9 | ~ spl92_24), inference(avatar_split_clause, [], [f600, f494, f434])).
fof(f600, plain, (~ p502(sK91) | ~ p202(sK91)), inference(resolution, [], [f277, f523])).
fof(f277, plain, ! [X0] : (~ sP40(X0) | ~ p502(X0) | ~ p202(X0)), inference(cnf_transformation, [], [f51])).
fof(f11795, plain, (~ spl92_277 | ~ spl92_280), inference(avatar_contradiction_clause, [], [f11794])).
fof(f11794, plain, ($false | (~ spl92_277 | ~ spl92_280)), inference(subsumption_resolution, [], [f11793, f544])).
fof(f544, plain, sP6(sK91), inference(resolution, [], [f523, f249])).
fof(f249, plain, ! [X0] : (~ sP40(X0) | sP6(X0)), inference(cnf_transformation, [], [f51])).
fof(f11793, plain, (~ sP6(sK91) | (~ spl92_277 | ~ spl92_280)), inference(subsumption_resolution, [], [f11792, f3054])).
fof(f3054, plain, (p304(sK78(sK91)) | ~ spl92_277), inference(avatar_component_clause, [], [f3053])).
fof(f3053, plain, (spl92_277 <=> p304(sK78(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_277])])).
fof(f11792, plain, (~ p304(sK78(sK91)) | ~ sP6(sK91) | ~ spl92_280), inference(resolution, [], [f3073, f375])).
fof(f375, plain, ! [X0] : (~ p204(sK77(X0)) | ~ p304(sK78(X0)) | ~ sP6(X0)), inference(cnf_transformation, [], [f191])).
fof(f191, plain, ! [X0] : ((~ p204(sK77(X0)) & r1(X0, sK77(X0))) | (~ p304(sK78(X0)) & r1(X0, sK78(X0))) | ~ sP6(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK77, sK78])], [f188, f190, f189])).
fof(f189, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) => (~ p204(sK77(X0)) & r1(X0, sK77(X0)))), introduced(choice_axiom, [])).
fof(f190, plain, ! [X0] : (? [X2] : (~ p304(X2) & r1(X0, X2)) => (~ p304(sK78(X0)) & r1(X0, sK78(X0)))), introduced(choice_axiom, [])).
fof(f188, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) | ? [X2] : (~ p304(X2) & r1(X0, X2)) | ~ sP6(X0)), inference(rectify, [], [f187])).
fof(f187, plain, ! [X11] : (? [X34] : (~ p204(X34) & r1(X11, X34)) | ? [X35] : (~ p304(X35) & r1(X11, X35)) | ~ sP6(X11)), inference(nnf_transformation, [], [f14])).
fof(f3073, plain, (p204(sK77(sK91)) | ~ spl92_280), inference(avatar_component_clause, [], [f3071])).
fof(f3071, plain, (spl92_280 <=> p204(sK77(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_280])])).
fof(f11712, plain, (~ spl92_280 | spl92_277 | ~ spl92_12), inference(avatar_split_clause, [], [f11711, f446, f3053, f3071])).
fof(f446, plain, (spl92_12 <=> ! [X8] : (p304(X8) | ~ r1(sK91, X8))), introduced(avatar_definition, [new_symbols(naming, [spl92_12])])).
fof(f11711, plain, (p304(sK78(sK91)) | ~ p204(sK77(sK91)) | ~ spl92_12), inference(subsumption_resolution, [], [f7390, f544])).
fof(f7390, plain, (p304(sK78(sK91)) | ~ p204(sK77(sK91)) | ~ sP6(sK91) | ~ spl92_12), inference(resolution, [], [f447, f374])).
fof(f374, plain, ! [X0] : (r1(X0, sK78(X0)) | ~ p204(sK77(X0)) | ~ sP6(X0)), inference(cnf_transformation, [], [f191])).
fof(f447, plain, (! [X8] : (~ r1(sK91, X8) | p304(X8)) | ~ spl92_12), inference(avatar_component_clause, [], [f446])).
fof(f11706, plain, (~ spl92_454 | ~ spl92_455), inference(avatar_contradiction_clause, [], [f11705])).
fof(f11705, plain, ($false | (~ spl92_454 | ~ spl92_455)), inference(subsumption_resolution, [], [f11704, f531])).
fof(f531, plain, sP1(sK91), inference(resolution, [], [f523, f233])).
fof(f233, plain, ! [X0] : (~ sP40(X0) | sP1(X0)), inference(cnf_transformation, [], [f51])).
fof(f11704, plain, (~ sP1(sK91) | (~ spl92_454 | ~ spl92_455)), inference(subsumption_resolution, [], [f11703, f4538])).
fof(f4538, plain, (p405(sK88(sK91)) | ~ spl92_454), inference(avatar_component_clause, [], [f4537])).
fof(f4537, plain, (spl92_454 <=> p405(sK88(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_454])])).
fof(f11703, plain, (~ p405(sK88(sK91)) | ~ sP1(sK91) | ~ spl92_455), inference(resolution, [], [f4545, f395])).
fof(f395, plain, ! [X0] : (~ p205(sK87(X0)) | ~ p405(sK88(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f216])).
fof(f216, plain, ! [X0] : ((~ p205(sK87(X0)) & r1(X0, sK87(X0))) | (~ p405(sK88(X0)) & r1(X0, sK88(X0))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK87, sK88])], [f213, f215, f214])).
fof(f214, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) => (~ p205(sK87(X0)) & r1(X0, sK87(X0)))), introduced(choice_axiom, [])).
fof(f215, plain, ! [X0] : (? [X2] : (~ p405(X2) & r1(X0, X2)) => (~ p405(sK88(X0)) & r1(X0, sK88(X0)))), introduced(choice_axiom, [])).
fof(f213, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) | ? [X2] : (~ p405(X2) & r1(X0, X2)) | ~ sP1(X0)), inference(rectify, [], [f212])).
fof(f212, plain, ! [X11] : (? [X52] : (~ p205(X52) & r1(X11, X52)) | ? [X53] : (~ p405(X53) & r1(X11, X53)) | ~ sP1(X11)), inference(nnf_transformation, [], [f9])).
fof(f4545, plain, (p205(sK87(sK91)) | ~ spl92_455), inference(avatar_component_clause, [], [f4543])).
fof(f4543, plain, (spl92_455 <=> p205(sK87(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_455])])).
fof(f11698, plain, (~ spl92_455 | spl92_454 | ~ spl92_16), inference(avatar_split_clause, [], [f11697, f462, f4537, f4543])).
fof(f11697, plain, (p405(sK88(sK91)) | ~ p205(sK87(sK91)) | ~ spl92_16), inference(subsumption_resolution, [], [f10194, f531])).
fof(f10194, plain, (p405(sK88(sK91)) | ~ p205(sK87(sK91)) | ~ sP1(sK91) | ~ spl92_16), inference(resolution, [], [f463, f394])).
fof(f394, plain, ! [X0] : (r1(X0, sK88(X0)) | ~ p205(sK87(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f216])).
fof(f11684, plain, (~ spl92_454 | spl92_455 | ~ spl92_6), inference(avatar_split_clause, [], [f11683, f425, f4543, f4537])).
fof(f11683, plain, (p205(sK87(sK91)) | ~ p405(sK88(sK91)) | ~ spl92_6), inference(subsumption_resolution, [], [f9289, f531])).
fof(f9289, plain, (p205(sK87(sK91)) | ~ p405(sK88(sK91)) | ~ sP1(sK91) | ~ spl92_6), inference(resolution, [], [f426, f393])).
fof(f393, plain, ! [X0] : (r1(X0, sK87(X0)) | ~ p405(sK88(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f216])).
fof(f11676, plain, (spl92_454 | ~ spl92_16 | spl92_560), inference(avatar_split_clause, [], [f11675, f5433, f462, f4537])).
fof(f5433, plain, (spl92_560 <=> r1(sK91, sK87(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_560])])).
fof(f11675, plain, (p405(sK88(sK91)) | (~ spl92_16 | spl92_560)), inference(subsumption_resolution, [], [f11674, f531])).
fof(f11674, plain, (p405(sK88(sK91)) | ~ sP1(sK91) | (~ spl92_16 | spl92_560)), inference(subsumption_resolution, [], [f10193, f5434])).
fof(f5434, plain, (~ r1(sK91, sK87(sK91)) | spl92_560), inference(avatar_component_clause, [], [f5433])).
fof(f10193, plain, (p405(sK88(sK91)) | r1(sK91, sK87(sK91)) | ~ sP1(sK91) | ~ spl92_16), inference(resolution, [], [f463, f392])).
fof(f392, plain, ! [X0] : (r1(X0, sK88(X0)) | r1(X0, sK87(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f216])).
fof(f11655, plain, (~ spl92_6 | spl92_455 | ~ spl92_560), inference(avatar_contradiction_clause, [], [f11654])).
fof(f11654, plain, ($false | (~ spl92_6 | spl92_455 | ~ spl92_560)), inference(subsumption_resolution, [], [f11650, f4544])).
fof(f4544, plain, (~ p205(sK87(sK91)) | spl92_455), inference(avatar_component_clause, [], [f4543])).
fof(f11650, plain, (p205(sK87(sK91)) | (~ spl92_6 | ~ spl92_560)), inference(resolution, [], [f5435, f426])).
fof(f5435, plain, (r1(sK91, sK87(sK91)) | ~ spl92_560), inference(avatar_component_clause, [], [f5433])).
fof(f11235, plain, (~ spl92_6 | spl92_427 | ~ spl92_539), inference(avatar_contradiction_clause, [], [f11234])).
fof(f11234, plain, ($false | (~ spl92_6 | spl92_427 | ~ spl92_539)), inference(subsumption_resolution, [], [f11230, f4297])).
fof(f4297, plain, (~ p205(sK85(sK91)) | spl92_427), inference(avatar_component_clause, [], [f4296])).
fof(f11230, plain, (p205(sK85(sK91)) | (~ spl92_6 | ~ spl92_539)), inference(resolution, [], [f5232, f426])).
fof(f5232, plain, (r1(sK91, sK85(sK91)) | ~ spl92_539), inference(avatar_component_clause, [], [f5230])).
fof(f10819, plain, (spl92_376 | ~ spl92_16 | spl92_518), inference(avatar_split_clause, [], [f10818, f5027, f462, f3846])).
fof(f10818, plain, (p405(sK84(sK91)) | (~ spl92_16 | spl92_518)), inference(subsumption_resolution, [], [f10817, f535])).
fof(f10817, plain, (p405(sK84(sK91)) | ~ sP3(sK91) | (~ spl92_16 | spl92_518)), inference(subsumption_resolution, [], [f10187, f5028])).
fof(f5028, plain, (~ r1(sK91, sK83(sK91)) | spl92_518), inference(avatar_component_clause, [], [f5027])).
fof(f10187, plain, (p405(sK84(sK91)) | r1(sK91, sK83(sK91)) | ~ sP3(sK91) | ~ spl92_16), inference(resolution, [], [f463, f384])).
fof(f384, plain, ! [X0] : (r1(X0, sK84(X0)) | r1(X0, sK83(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f206])).
fof(f10133, plain, (~ spl92_17 | ~ spl92_321), inference(avatar_split_clause, [], [f10132, f3341, f465])).
fof(f3341, plain, (spl92_321 <=> p304(sK60(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_321])])).
fof(f10132, plain, (~ p404(sK91) | ~ spl92_321), inference(subsumption_resolution, [], [f10128, f540])).
fof(f540, plain, sP20(sK91), inference(resolution, [], [f523, f245])).
fof(f245, plain, ! [X0] : (~ sP40(X0) | sP20(X0)), inference(cnf_transformation, [], [f51])).
fof(f10128, plain, (~ p404(sK91) | ~ sP20(sK91) | ~ spl92_321), inference(resolution, [], [f3343, f339])).
fof(f339, plain, ! [X0] : (~ p304(sK60(X0)) | ~ p404(X0) | ~ sP20(X0)), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ! [X0] : ((~ p304(sK60(X0)) & r1(X0, sK60(X0))) | ~ p404(X0) | ~ sP20(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK60])], [f129, f130])).
fof(f130, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) => (~ p304(sK60(X0)) & r1(X0, sK60(X0)))), introduced(choice_axiom, [])).
fof(f129, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) | ~ p404(X0) | ~ sP20(X0)), inference(rectify, [], [f128])).
fof(f128, plain, ! [X11] : (? [X39] : (~ p304(X39) & r1(X11, X39)) | ~ p404(X11) | ~ sP20(X11)), inference(nnf_transformation, [], [f28])).
fof(f3343, plain, (p304(sK60(sK91)) | ~ spl92_321), inference(avatar_component_clause, [], [f3341])).
fof(f9937, plain, (~ spl92_18 | ~ spl92_311), inference(avatar_split_clause, [], [f9936, f3281, f469])).
fof(f3281, plain, (spl92_311 <=> p103(sK47(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_311])])).
fof(f9936, plain, (~ p403(sK91) | ~ spl92_311), inference(subsumption_resolution, [], [f9902, f556])).
fof(f556, plain, sP33(sK91), inference(resolution, [], [f523, f267])).
fof(f267, plain, ! [X0] : (~ sP40(X0) | sP33(X0)), inference(cnf_transformation, [], [f51])).
fof(f9902, plain, (~ p403(sK91) | ~ sP33(sK91) | ~ spl92_311), inference(resolution, [], [f3283, f313])).
fof(f313, plain, ! [X0] : (~ p103(sK47(X0)) | ~ p403(X0) | ~ sP33(X0)), inference(cnf_transformation, [], [f79])).
fof(f79, plain, ! [X0] : ((~ p103(sK47(X0)) & r1(X0, sK47(X0))) | ~ p403(X0) | ~ sP33(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK47])], [f77, f78])).
fof(f78, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) => (~ p103(sK47(X0)) & r1(X0, sK47(X0)))), introduced(choice_axiom, [])).
fof(f77, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) | ~ p403(X0) | ~ sP33(X0)), inference(rectify, [], [f76])).
fof(f76, plain, ! [X11] : (? [X20] : (~ p103(X20) & r1(X11, X20)) | ~ p403(X11) | ~ sP33(X11)), inference(nnf_transformation, [], [f41])).
fof(f3283, plain, (p103(sK47(sK91)) | ~ spl92_311), inference(avatar_component_clause, [], [f3281])).
fof(f9925, plain, (~ spl92_17 | spl92_321 | ~ spl92_12), inference(avatar_split_clause, [], [f9924, f446, f3341, f465])).
fof(f9924, plain, (p304(sK60(sK91)) | ~ p404(sK91) | ~ spl92_12), inference(subsumption_resolution, [], [f7368, f540])).
fof(f7368, plain, (p304(sK60(sK91)) | ~ p404(sK91) | ~ sP20(sK91) | ~ spl92_12), inference(resolution, [], [f447, f338])).
fof(f338, plain, ! [X0] : (r1(X0, sK60(X0)) | ~ p404(X0) | ~ sP20(X0)), inference(cnf_transformation, [], [f131])).
fof(f9515, plain, (~ spl92_22 | ~ spl92_194), inference(avatar_split_clause, [], [f9514, f2345, f486])).
fof(f2345, plain, (spl92_194 <=> p304(sK61(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_194])])).
fof(f9514, plain, (~ p504(sK91) | ~ spl92_194), inference(subsumption_resolution, [], [f9490, f539])).
fof(f539, plain, sP19(sK91), inference(resolution, [], [f523, f244])).
fof(f244, plain, ! [X0] : (~ sP40(X0) | sP19(X0)), inference(cnf_transformation, [], [f51])).
fof(f9490, plain, (~ p504(sK91) | ~ sP19(sK91) | ~ spl92_194), inference(resolution, [], [f2347, f341])).
fof(f341, plain, ! [X0] : (~ p304(sK61(X0)) | ~ p504(X0) | ~ sP19(X0)), inference(cnf_transformation, [], [f135])).
fof(f135, plain, ! [X0] : ((~ p304(sK61(X0)) & r1(X0, sK61(X0))) | ~ p504(X0) | ~ sP19(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK61])], [f133, f134])).
fof(f134, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) => (~ p304(sK61(X0)) & r1(X0, sK61(X0)))), introduced(choice_axiom, [])).
fof(f133, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) | ~ p504(X0) | ~ sP19(X0)), inference(rectify, [], [f132])).
fof(f132, plain, ! [X11] : (? [X40] : (~ p304(X40) & r1(X11, X40)) | ~ p504(X11) | ~ sP19(X11)), inference(nnf_transformation, [], [f27])).
fof(f2347, plain, (p304(sK61(sK91)) | ~ spl92_194), inference(avatar_component_clause, [], [f2345])).
fof(f9513, plain, (~ spl92_18 | spl92_311 | ~ spl92_3), inference(avatar_split_clause, [], [f9512, f414, f3281, f469])).
fof(f9512, plain, (p103(sK47(sK91)) | ~ p403(sK91) | ~ spl92_3), inference(subsumption_resolution, [], [f7286, f556])).
fof(f7286, plain, (p103(sK47(sK91)) | ~ p403(sK91) | ~ sP33(sK91) | ~ spl92_3), inference(resolution, [], [f415, f312])).
fof(f312, plain, ! [X0] : (r1(X0, sK47(X0)) | ~ p403(X0) | ~ sP33(X0)), inference(cnf_transformation, [], [f79])).
fof(f9220, plain, (~ spl92_22 | spl92_194 | ~ spl92_12), inference(avatar_split_clause, [], [f9219, f446, f2345, f486])).
fof(f9219, plain, (p304(sK61(sK91)) | ~ p504(sK91) | ~ spl92_12), inference(subsumption_resolution, [], [f7369, f539])).
fof(f7369, plain, (p304(sK61(sK91)) | ~ p504(sK91) | ~ sP19(sK91) | ~ spl92_12), inference(resolution, [], [f447, f340])).
fof(f340, plain, ! [X0] : (r1(X0, sK61(X0)) | ~ p504(X0) | ~ sP19(X0)), inference(cnf_transformation, [], [f135])).
fof(f9216, plain, (spl92_513 | ~ spl92_12 | spl92_277), inference(avatar_split_clause, [], [f9215, f3053, f446, f4984])).
fof(f4984, plain, (spl92_513 <=> r1(sK91, sK77(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_513])])).
fof(f9215, plain, (r1(sK91, sK77(sK91)) | (~ spl92_12 | spl92_277)), inference(subsumption_resolution, [], [f9214, f544])).
fof(f9214, plain, (r1(sK91, sK77(sK91)) | ~ sP6(sK91) | (~ spl92_12 | spl92_277)), inference(subsumption_resolution, [], [f7389, f3055])).
fof(f3055, plain, (~ p304(sK78(sK91)) | spl92_277), inference(avatar_component_clause, [], [f3053])).
fof(f7389, plain, (p304(sK78(sK91)) | r1(sK91, sK77(sK91)) | ~ sP6(sK91) | ~ spl92_12), inference(resolution, [], [f447, f372])).
fof(f372, plain, ! [X0] : (r1(X0, sK78(X0)) | r1(X0, sK77(X0)) | ~ sP6(X0)), inference(cnf_transformation, [], [f191])).
fof(f9208, plain, (~ spl92_7 | spl92_280 | ~ spl92_513), inference(avatar_contradiction_clause, [], [f9207])).
fof(f9207, plain, ($false | (~ spl92_7 | spl92_280 | ~ spl92_513)), inference(subsumption_resolution, [], [f9203, f3072])).
fof(f3072, plain, (~ p204(sK77(sK91)) | spl92_280), inference(avatar_component_clause, [], [f3071])).
fof(f9203, plain, (p204(sK77(sK91)) | (~ spl92_7 | ~ spl92_513)), inference(resolution, [], [f4986, f429])).
fof(f4986, plain, (r1(sK91, sK77(sK91)) | ~ spl92_513), inference(avatar_component_clause, [], [f4984])).
fof(f8998, plain, (spl92_266 | ~ spl92_12 | spl92_480), inference(avatar_split_clause, [], [f8997, f4713, f446, f2978])).
fof(f8997, plain, (p304(sK76(sK91)) | (~ spl92_12 | spl92_480)), inference(subsumption_resolution, [], [f8996, f548])).
fof(f8996, plain, (p304(sK76(sK91)) | ~ sP7(sK91) | (~ spl92_12 | spl92_480)), inference(subsumption_resolution, [], [f7386, f4714])).
fof(f4714, plain, (~ r1(sK91, sK75(sK91)) | spl92_480), inference(avatar_component_clause, [], [f4713])).
fof(f7386, plain, (p304(sK76(sK91)) | r1(sK91, sK75(sK91)) | ~ sP7(sK91) | ~ spl92_12), inference(resolution, [], [f447, f368])).
fof(f368, plain, ! [X0] : (r1(X0, sK76(X0)) | r1(X0, sK75(X0)) | ~ sP7(X0)), inference(cnf_transformation, [], [f186])).
fof(f8336, plain, (~ spl92_185 | ~ spl92_184), inference(avatar_split_clause, [], [f8335, f2283, f2287])).
fof(f2287, plain, (spl92_185 <=> p203(sK72(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_185])])).
fof(f2283, plain, (spl92_184 <=> p103(sK71(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_184])])).
fof(f8335, plain, (~ p203(sK72(sK91)) | ~ spl92_184), inference(subsumption_resolution, [], [f8329, f558])).
fof(f558, plain, sP9(sK91), inference(resolution, [], [f523, f269])).
fof(f269, plain, ! [X0] : (~ sP40(X0) | sP9(X0)), inference(cnf_transformation, [], [f51])).
fof(f8329, plain, (~ p203(sK72(sK91)) | ~ sP9(sK91) | ~ spl92_184), inference(resolution, [], [f2285, f363])).
fof(f363, plain, ! [X0] : (~ p103(sK71(X0)) | ~ p203(sK72(X0)) | ~ sP9(X0)), inference(cnf_transformation, [], [f176])).
fof(f176, plain, ! [X0] : ((~ p103(sK71(X0)) & r1(X0, sK71(X0))) | (~ p203(sK72(X0)) & r1(X0, sK72(X0))) | ~ sP9(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK71, sK72])], [f173, f175, f174])).
fof(f174, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) => (~ p103(sK71(X0)) & r1(X0, sK71(X0)))), introduced(choice_axiom, [])).
fof(f175, plain, ! [X0] : (? [X2] : (~ p203(X2) & r1(X0, X2)) => (~ p203(sK72(X0)) & r1(X0, sK72(X0)))), introduced(choice_axiom, [])).
fof(f173, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) | ? [X2] : (~ p203(X2) & r1(X0, X2)) | ~ sP9(X0)), inference(rectify, [], [f172])).
fof(f172, plain, ! [X11] : (? [X17] : (~ p103(X17) & r1(X11, X17)) | ? [X18] : (~ p203(X18) & r1(X11, X18)) | ~ sP9(X11)), inference(nnf_transformation, [], [f17])).
fof(f2285, plain, (p103(sK71(sK91)) | ~ spl92_184), inference(avatar_component_clause, [], [f2283])).
fof(f8334, plain, (~ spl92_10 | ~ spl92_20), inference(avatar_split_clause, [], [f8333, f477, f438])).
fof(f8333, plain, (~ p201(sK91) | ~ spl92_20), inference(subsumption_resolution, [], [f615, f479])).
fof(f479, plain, (p401(sK91) | ~ spl92_20), inference(avatar_component_clause, [], [f477])).
fof(f615, plain, (~ p401(sK91) | ~ p201(sK91)), inference(resolution, [], [f293, f523])).
fof(f293, plain, ! [X0] : (~ sP40(X0) | ~ p401(X0) | ~ p201(X0)), inference(cnf_transformation, [], [f51])).
fof(f8322, plain, (spl92_185 | ~ spl92_8 | ~ spl92_184), inference(avatar_split_clause, [], [f8321, f2283, f431, f2287])).
fof(f8321, plain, (p203(sK72(sK91)) | (~ spl92_8 | ~ spl92_184)), inference(subsumption_resolution, [], [f8320, f558])).
fof(f8320, plain, (p203(sK72(sK91)) | ~ sP9(sK91) | (~ spl92_8 | ~ spl92_184)), inference(subsumption_resolution, [], [f8070, f2285])).
fof(f8070, plain, (p203(sK72(sK91)) | ~ p103(sK71(sK91)) | ~ sP9(sK91) | ~ spl92_8), inference(resolution, [], [f432, f362])).
fof(f362, plain, ! [X0] : (r1(X0, sK72(X0)) | ~ p103(sK71(X0)) | ~ sP9(X0)), inference(cnf_transformation, [], [f176])).
fof(f8308, plain, (spl92_472 | ~ spl92_8 | spl92_185), inference(avatar_split_clause, [], [f8307, f2287, f431, f4657])).
fof(f4657, plain, (spl92_472 <=> r1(sK91, sK71(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_472])])).
fof(f8307, plain, (r1(sK91, sK71(sK91)) | (~ spl92_8 | spl92_185)), inference(subsumption_resolution, [], [f8306, f558])).
fof(f8306, plain, (r1(sK91, sK71(sK91)) | ~ sP9(sK91) | (~ spl92_8 | spl92_185)), inference(subsumption_resolution, [], [f8069, f2289])).
fof(f2289, plain, (~ p203(sK72(sK91)) | spl92_185), inference(avatar_component_clause, [], [f2287])).
fof(f8069, plain, (p203(sK72(sK91)) | r1(sK91, sK71(sK91)) | ~ sP9(sK91) | ~ spl92_8), inference(resolution, [], [f432, f360])).
fof(f360, plain, ! [X0] : (r1(X0, sK72(X0)) | r1(X0, sK71(X0)) | ~ sP9(X0)), inference(cnf_transformation, [], [f176])).
fof(f8304, plain, (~ spl92_3 | spl92_184 | ~ spl92_472), inference(avatar_contradiction_clause, [], [f8303])).
fof(f8303, plain, ($false | (~ spl92_3 | spl92_184 | ~ spl92_472)), inference(subsumption_resolution, [], [f8300, f2284])).
fof(f2284, plain, (~ p103(sK71(sK91)) | spl92_184), inference(avatar_component_clause, [], [f2283])).
fof(f8300, plain, (p103(sK71(sK91)) | (~ spl92_3 | ~ spl92_472)), inference(resolution, [], [f4659, f415])).
fof(f4659, plain, (r1(sK91, sK71(sK91)) | ~ spl92_472), inference(avatar_component_clause, [], [f4657])).
fof(f7988, plain, (~ spl92_3 | ~ spl92_28), inference(avatar_contradiction_clause, [], [f7987])).
fof(f7987, plain, ($false | (~ spl92_3 | ~ spl92_28)), inference(subsumption_resolution, [], [f7986, f554])).
fof(f554, plain, sP31(sK91), inference(resolution, [], [f523, f265])).
fof(f265, plain, ! [X0] : (~ sP40(X0) | sP31(X0)), inference(cnf_transformation, [], [f51])).
fof(f7986, plain, (~ sP31(sK91) | (~ spl92_3 | ~ spl92_28)), inference(subsumption_resolution, [], [f7985, f513])).
fof(f7985, plain, (~ p603(sK91) | ~ sP31(sK91) | (~ spl92_3 | ~ spl92_28)), inference(resolution, [], [f7343, f317])).
fof(f317, plain, ! [X0] : (~ p103(sK49(X0)) | ~ p603(X0) | ~ sP31(X0)), inference(cnf_transformation, [], [f87])).
fof(f87, plain, ! [X0] : ((~ p103(sK49(X0)) & r1(X0, sK49(X0))) | ~ p603(X0) | ~ sP31(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK49])], [f85, f86])).
fof(f86, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) => (~ p103(sK49(X0)) & r1(X0, sK49(X0)))), introduced(choice_axiom, [])).
fof(f85, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) | ~ p603(X0) | ~ sP31(X0)), inference(rectify, [], [f84])).
fof(f84, plain, ! [X11] : (? [X22] : (~ p103(X22) & r1(X11, X22)) | ~ p603(X11) | ~ sP31(X11)), inference(nnf_transformation, [], [f39])).
fof(f7343, plain, (p103(sK49(sK91)) | (~ spl92_3 | ~ spl92_28)), inference(subsumption_resolution, [], [f7342, f554])).
fof(f7342, plain, (p103(sK49(sK91)) | ~ sP31(sK91) | (~ spl92_3 | ~ spl92_28)), inference(subsumption_resolution, [], [f7288, f513])).
fof(f7288, plain, (p103(sK49(sK91)) | ~ p603(sK91) | ~ sP31(sK91) | ~ spl92_3), inference(resolution, [], [f415, f316])).
fof(f316, plain, ! [X0] : (r1(X0, sK49(X0)) | ~ p603(X0) | ~ sP31(X0)), inference(cnf_transformation, [], [f87])).
fof(f7270, plain, (~ spl92_4 | ~ spl92_9), inference(avatar_contradiction_clause, [], [f7269])).
fof(f7269, plain, ($false | (~ spl92_4 | ~ spl92_9)), inference(subsumption_resolution, [], [f7268, f563])).
fof(f563, plain, sP39(sK91), inference(resolution, [], [f523, f284])).
fof(f284, plain, ! [X0] : (~ sP40(X0) | sP39(X0)), inference(cnf_transformation, [], [f51])).
fof(f7268, plain, (~ sP39(sK91) | (~ spl92_4 | ~ spl92_9)), inference(subsumption_resolution, [], [f7267, f436])).
fof(f436, plain, (p202(sK91) | ~ spl92_9), inference(avatar_component_clause, [], [f434])).
fof(f7267, plain, (~ p202(sK91) | ~ sP39(sK91) | (~ spl92_4 | ~ spl92_9)), inference(resolution, [], [f7052, f301])).
fof(f301, plain, ! [X0] : (~ p102(sK41(X0)) | ~ p202(X0) | ~ sP39(X0)), inference(cnf_transformation, [], [f55])).
fof(f55, plain, ! [X0] : ((~ p102(sK41(X0)) & r1(X0, sK41(X0))) | ~ p202(X0) | ~ sP39(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK41])], [f53, f54])).
fof(f54, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) => (~ p102(sK41(X0)) & r1(X0, sK41(X0)))), introduced(choice_axiom, [])).
fof(f53, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) | ~ p202(X0) | ~ sP39(X0)), inference(rectify, [], [f52])).
fof(f52, plain, ! [X11] : (? [X12] : (~ p102(X12) & r1(X11, X12)) | ~ p202(X11) | ~ sP39(X11)), inference(nnf_transformation, [], [f47])).
fof(f7052, plain, (p102(sK41(sK91)) | (~ spl92_4 | ~ spl92_9)), inference(subsumption_resolution, [], [f7051, f563])).
fof(f7051, plain, (p102(sK41(sK91)) | ~ sP39(sK91) | (~ spl92_4 | ~ spl92_9)), inference(subsumption_resolution, [], [f6991, f436])).
fof(f6991, plain, (p102(sK41(sK91)) | ~ p202(sK91) | ~ sP39(sK91) | ~ spl92_4), inference(resolution, [], [f418, f300])).
fof(f300, plain, ! [X0] : (r1(X0, sK41(X0)) | ~ p202(X0) | ~ sP39(X0)), inference(cnf_transformation, [], [f55])).
fof(f6989, plain, (~ spl92_1 | ~ spl92_21), inference(avatar_contradiction_clause, [], [f6988])).
fof(f6988, plain, ($false | (~ spl92_1 | ~ spl92_21)), inference(subsumption_resolution, [], [f6987, f534])).
fof(f534, plain, sP17(sK91), inference(resolution, [], [f523, f236])).
fof(f236, plain, ! [X0] : (~ sP40(X0) | sP17(X0)), inference(cnf_transformation, [], [f51])).
fof(f6987, plain, (~ sP17(sK91) | (~ spl92_1 | ~ spl92_21)), inference(subsumption_resolution, [], [f6986, f484])).
fof(f6986, plain, (~ p505(sK91) | ~ sP17(sK91) | (~ spl92_1 | ~ spl92_21)), inference(resolution, [], [f6480, f345])).
fof(f345, plain, ! [X0] : (~ p105(sK63(X0)) | ~ p505(X0) | ~ sP17(X0)), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ! [X0] : ((~ p105(sK63(X0)) & r1(X0, sK63(X0))) | ~ p505(X0) | ~ sP17(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK63])], [f141, f142])).
fof(f142, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) => (~ p105(sK63(X0)) & r1(X0, sK63(X0)))), introduced(choice_axiom, [])).
fof(f141, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) | ~ p505(X0) | ~ sP17(X0)), inference(rectify, [], [f140])).
fof(f140, plain, ! [X11] : (? [X48] : (~ p105(X48) & r1(X11, X48)) | ~ p505(X11) | ~ sP17(X11)), inference(nnf_transformation, [], [f25])).
fof(f6480, plain, (p105(sK63(sK91)) | (~ spl92_1 | ~ spl92_21)), inference(subsumption_resolution, [], [f6479, f534])).
fof(f6479, plain, (p105(sK63(sK91)) | ~ sP17(sK91) | (~ spl92_1 | ~ spl92_21)), inference(subsumption_resolution, [], [f6431, f484])).
fof(f6431, plain, (p105(sK63(sK91)) | ~ p505(sK91) | ~ sP17(sK91) | ~ spl92_1), inference(resolution, [], [f409, f344])).
fof(f344, plain, ! [X0] : (r1(X0, sK63(X0)) | ~ p505(X0) | ~ sP17(X0)), inference(cnf_transformation, [], [f143])).
fof(f6496, plain, (~ spl92_6 | ~ spl92_21), inference(avatar_contradiction_clause, [], [f6495])).
fof(f6495, plain, ($false | (~ spl92_6 | ~ spl92_21)), inference(subsumption_resolution, [], [f6494, f530])).
fof(f530, plain, sP15(sK91), inference(resolution, [], [f523, f232])).
fof(f232, plain, ! [X0] : (~ sP40(X0) | sP15(X0)), inference(cnf_transformation, [], [f51])).
fof(f6494, plain, (~ sP15(sK91) | (~ spl92_6 | ~ spl92_21)), inference(subsumption_resolution, [], [f6493, f484])).
fof(f6493, plain, (~ p505(sK91) | ~ sP15(sK91) | (~ spl92_6 | ~ spl92_21)), inference(resolution, [], [f3774, f349])).
fof(f349, plain, ! [X0] : (~ p205(sK65(X0)) | ~ p505(X0) | ~ sP15(X0)), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ! [X0] : ((~ p205(sK65(X0)) & r1(X0, sK65(X0))) | ~ p505(X0) | ~ sP15(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK65])], [f149, f150])).
fof(f150, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) => (~ p205(sK65(X0)) & r1(X0, sK65(X0)))), introduced(choice_axiom, [])).
fof(f149, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) | ~ p505(X0) | ~ sP15(X0)), inference(rectify, [], [f148])).
fof(f148, plain, ! [X11] : (? [X54] : (~ p205(X54) & r1(X11, X54)) | ~ p505(X11) | ~ sP15(X11)), inference(nnf_transformation, [], [f23])).
fof(f3774, plain, (p205(sK65(sK91)) | (~ spl92_6 | ~ spl92_21)), inference(subsumption_resolution, [], [f3773, f530])).
fof(f3773, plain, (p205(sK65(sK91)) | ~ sP15(sK91) | (~ spl92_6 | ~ spl92_21)), inference(subsumption_resolution, [], [f3747, f484])).
fof(f3747, plain, (p205(sK65(sK91)) | ~ p505(sK91) | ~ sP15(sK91) | ~ spl92_6), inference(resolution, [], [f426, f348])).
fof(f348, plain, ! [X0] : (r1(X0, sK65(X0)) | ~ p505(X0) | ~ sP15(X0)), inference(cnf_transformation, [], [f151])).
fof(f6407, plain, (~ spl92_2 | ~ spl92_27), inference(avatar_contradiction_clause, [], [f6406])).
fof(f6406, plain, ($false | (~ spl92_2 | ~ spl92_27)), inference(subsumption_resolution, [], [f6405, f545])).
fof(f545, plain, sP24(sK91), inference(resolution, [], [f523, f250])).
fof(f250, plain, ! [X0] : (~ sP40(X0) | sP24(X0)), inference(cnf_transformation, [], [f51])).
fof(f6405, plain, (~ sP24(sK91) | (~ spl92_2 | ~ spl92_27)), inference(subsumption_resolution, [], [f6404, f509])).
fof(f6404, plain, (~ p604(sK91) | ~ sP24(sK91) | (~ spl92_2 | ~ spl92_27)), inference(resolution, [], [f3709, f331])).
fof(f331, plain, ! [X0] : (~ p104(sK56(X0)) | ~ p604(X0) | ~ sP24(X0)), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ! [X0] : ((~ p104(sK56(X0)) & r1(X0, sK56(X0))) | ~ p604(X0) | ~ sP24(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK56])], [f113, f114])).
fof(f114, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) => (~ p104(sK56(X0)) & r1(X0, sK56(X0)))), introduced(choice_axiom, [])).
fof(f113, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) | ~ p604(X0) | ~ sP24(X0)), inference(rectify, [], [f112])).
fof(f112, plain, ! [X11] : (? [X33] : (~ p104(X33) & r1(X11, X33)) | ~ p604(X11) | ~ sP24(X11)), inference(nnf_transformation, [], [f32])).
fof(f3709, plain, (p104(sK56(sK91)) | (~ spl92_2 | ~ spl92_27)), inference(subsumption_resolution, [], [f3708, f545])).
fof(f3708, plain, (p104(sK56(sK91)) | ~ sP24(sK91) | (~ spl92_2 | ~ spl92_27)), inference(subsumption_resolution, [], [f3681, f509])).
fof(f3681, plain, (p104(sK56(sK91)) | ~ p604(sK91) | ~ sP24(sK91) | ~ spl92_2), inference(resolution, [], [f412, f330])).
fof(f330, plain, ! [X0] : (r1(X0, sK56(X0)) | ~ p604(X0) | ~ sP24(X0)), inference(cnf_transformation, [], [f115])).
fof(f4508, plain, (~ spl92_427 | ~ spl92_11 | spl92_426), inference(avatar_split_clause, [], [f4507, f4290, f443, f4296])).
fof(f4507, plain, (~ p205(sK85(sK91)) | (~ spl92_11 | spl92_426)), inference(subsumption_resolution, [], [f4506, f4292])).
fof(f4506, plain, (~ p205(sK85(sK91)) | p305(sK86(sK91)) | ~ spl92_11), inference(subsumption_resolution, [], [f4502, f532])).
fof(f4502, plain, (~ p205(sK85(sK91)) | ~ sP2(sK91) | p305(sK86(sK91)) | ~ spl92_11), inference(resolution, [], [f390, f444])).
fof(f390, plain, ! [X0] : (r1(X0, sK86(X0)) | ~ p205(sK85(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f211])).
fof(f4299, plain, (spl92_427 | ~ spl92_426 | ~ spl92_6), inference(avatar_split_clause, [], [f4294, f425, f4290, f4296])).
fof(f4294, plain, (~ p305(sK86(sK91)) | p205(sK85(sK91)) | ~ spl92_6), inference(subsumption_resolution, [], [f4281, f532])).
fof(f4281, plain, (~ p305(sK86(sK91)) | ~ sP2(sK91) | p205(sK85(sK91)) | ~ spl92_6), inference(resolution, [], [f389, f426])).
fof(f389, plain, ! [X0] : (r1(X0, sK85(X0)) | ~ p305(sK86(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f211])).
fof(f3646, plain, (~ spl92_27 | spl92_371 | ~ spl92_7), inference(avatar_split_clause, [], [f3641, f428, f3643, f507])).
fof(f3641, plain, (p204(sK59(sK91)) | ~ p604(sK91) | ~ spl92_7), inference(subsumption_resolution, [], [f985, f541])).
fof(f985, plain, (p204(sK59(sK91)) | ~ p604(sK91) | ~ sP21(sK91) | ~ spl92_7), inference(resolution, [], [f429, f336])).
fof(f3628, plain, (spl92_368 | ~ spl92_27 | ~ spl92_12), inference(avatar_split_clause, [], [f3623, f446, f507, f3625])).
fof(f3623, plain, (~ p604(sK91) | p304(sK62(sK91)) | ~ spl92_12), inference(subsumption_resolution, [], [f1402, f538])).
fof(f1402, plain, (~ p604(sK91) | ~ sP18(sK91) | p304(sK62(sK91)) | ~ spl92_12), inference(resolution, [], [f342, f447])).
fof(f342, plain, ! [X0] : (r1(X0, sK62(X0)) | ~ p604(X0) | ~ sP18(X0)), inference(cnf_transformation, [], [f139])).
fof(f3592, plain, (~ spl92_21 | ~ spl92_26), inference(avatar_split_clause, [], [f568, f503, f482])).
fof(f568, plain, (~ p605(sK91) | ~ p505(sK91)), inference(resolution, [], [f225, f523])).
fof(f225, plain, ! [X0] : (~ sP40(X0) | ~ p605(X0) | ~ p505(X0)), inference(cnf_transformation, [], [f51])).
fof(f3495, plain, (~ spl92_20 | ~ spl92_25), inference(avatar_split_clause, [], [f607, f498, f477])).
fof(f607, plain, (~ p501(sK91) | ~ p401(sK91)), inference(resolution, [], [f287, f523])).
fof(f287, plain, ! [X0] : (~ sP40(X0) | ~ p501(X0) | ~ p401(X0)), inference(cnf_transformation, [], [f51])).
fof(f3362, plain, (~ spl92_17 | spl92_324 | ~ spl92_7), inference(avatar_split_clause, [], [f3357, f428, f3359, f465])).
fof(f3357, plain, (p204(sK57(sK91)) | ~ p404(sK91) | ~ spl92_7), inference(subsumption_resolution, [], [f983, f543])).
fof(f983, plain, (p204(sK57(sK91)) | ~ p404(sK91) | ~ sP23(sK91) | ~ spl92_7), inference(resolution, [], [f429, f332])).
fof(f3344, plain, (spl92_321 | ~ spl92_17 | ~ spl92_12), inference(avatar_split_clause, [], [f3339, f446, f465, f3341])).
fof(f3339, plain, (~ p404(sK91) | p304(sK60(sK91)) | ~ spl92_12), inference(subsumption_resolution, [], [f1014, f540])).
fof(f1014, plain, (~ p404(sK91) | ~ sP20(sK91) | p304(sK60(sK91)) | ~ spl92_12), inference(resolution, [], [f338, f447])).
fof(f3272, plain, (~ spl92_16 | ~ spl92_26), inference(avatar_contradiction_clause, [], [f3271])).
fof(f3271, plain, ($false | (~ spl92_16 | ~ spl92_26)), inference(subsumption_resolution, [], [f3270, f524])).
fof(f524, plain, sP10(sK91), inference(resolution, [], [f523, f226])).
fof(f226, plain, ! [X0] : (~ sP40(X0) | sP10(X0)), inference(cnf_transformation, [], [f51])).
fof(f3270, plain, (~ sP10(sK91) | (~ spl92_16 | ~ spl92_26)), inference(subsumption_resolution, [], [f3269, f505])).
fof(f3269, plain, (~ p605(sK91) | ~ sP10(sK91) | (~ spl92_16 | ~ spl92_26)), inference(resolution, [], [f2666, f359])).
fof(f359, plain, ! [X0] : (~ p405(sK70(X0)) | ~ p605(X0) | ~ sP10(X0)), inference(cnf_transformation, [], [f171])).
fof(f171, plain, ! [X0] : ((~ p405(sK70(X0)) & r1(X0, sK70(X0))) | ~ p605(X0) | ~ sP10(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK70])], [f169, f170])).
fof(f170, plain, ! [X0] : (? [X1] : (~ p405(X1) & r1(X0, X1)) => (~ p405(sK70(X0)) & r1(X0, sK70(X0)))), introduced(choice_axiom, [])).
fof(f169, plain, ! [X0] : (? [X1] : (~ p405(X1) & r1(X0, X1)) | ~ p605(X0) | ~ sP10(X0)), inference(rectify, [], [f168])).
fof(f168, plain, ! [X11] : (? [X61] : (~ p405(X61) & r1(X11, X61)) | ~ p605(X11) | ~ sP10(X11)), inference(nnf_transformation, [], [f18])).
fof(f2666, plain, (p405(sK70(sK91)) | (~ spl92_16 | ~ spl92_26)), inference(subsumption_resolution, [], [f2665, f524])).
fof(f2665, plain, (p405(sK70(sK91)) | ~ sP10(sK91) | (~ spl92_16 | ~ spl92_26)), inference(subsumption_resolution, [], [f2657, f505])).
fof(f2657, plain, (p405(sK70(sK91)) | ~ p605(sK91) | ~ sP10(sK91) | ~ spl92_16), inference(resolution, [], [f463, f358])).
fof(f358, plain, ! [X0] : (r1(X0, sK70(X0)) | ~ p605(X0) | ~ sP10(X0)), inference(cnf_transformation, [], [f171])).
fof(f3074, plain, (spl92_280 | ~ spl92_277 | ~ spl92_7), inference(avatar_split_clause, [], [f3069, f428, f3053, f3071])).
fof(f3069, plain, (~ p304(sK78(sK91)) | p204(sK77(sK91)) | ~ spl92_7), inference(subsumption_resolution, [], [f3045, f544])).
fof(f3045, plain, (~ p304(sK78(sK91)) | ~ sP6(sK91) | p204(sK77(sK91)) | ~ spl92_7), inference(resolution, [], [f373, f429])).
fof(f373, plain, ! [X0] : (r1(X0, sK77(X0)) | ~ p304(sK78(X0)) | ~ sP6(X0)), inference(cnf_transformation, [], [f191])).
fof(f3029, plain, (~ spl92_272 | ~ spl92_12 | spl92_266), inference(avatar_split_clause, [], [f3028, f2978, f446, f3017])).
fof(f3028, plain, (~ p104(sK75(sK91)) | (~ spl92_12 | spl92_266)), inference(subsumption_resolution, [], [f3027, f2980])).
fof(f2980, plain, (~ p304(sK76(sK91)) | spl92_266), inference(avatar_component_clause, [], [f2978])).
fof(f3027, plain, (~ p104(sK75(sK91)) | p304(sK76(sK91)) | ~ spl92_12), inference(subsumption_resolution, [], [f3008, f548])).
fof(f3008, plain, (~ p104(sK75(sK91)) | ~ sP7(sK91) | p304(sK76(sK91)) | ~ spl92_12), inference(resolution, [], [f370, f447])).
fof(f370, plain, ! [X0] : (r1(X0, sK76(X0)) | ~ p104(sK75(X0)) | ~ sP7(X0)), inference(cnf_transformation, [], [f186])).
fof(f2960, plain, (~ spl92_261 | ~ spl92_7 | spl92_255), inference(avatar_split_clause, [], [f2959, f2903, f428, f2942])).
fof(f2959, plain, (~ p104(sK73(sK91)) | (~ spl92_7 | spl92_255)), inference(subsumption_resolution, [], [f2958, f2905])).
fof(f2958, plain, (~ p104(sK73(sK91)) | p204(sK74(sK91)) | ~ spl92_7), inference(subsumption_resolution, [], [f2934, f549])).
fof(f2934, plain, (~ p104(sK73(sK91)) | ~ sP8(sK91) | p204(sK74(sK91)) | ~ spl92_7), inference(resolution, [], [f366, f429])).
fof(f366, plain, ! [X0] : (r1(X0, sK74(X0)) | ~ p104(sK73(X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f181])).
fof(f2366, plain, (~ spl92_22 | spl92_197 | ~ spl92_7), inference(avatar_split_clause, [], [f2361, f428, f2363, f486])).
fof(f2361, plain, (p204(sK58(sK91)) | ~ p504(sK91) | ~ spl92_7), inference(subsumption_resolution, [], [f984, f542])).
fof(f984, plain, (p204(sK58(sK91)) | ~ p504(sK91) | ~ sP22(sK91) | ~ spl92_7), inference(resolution, [], [f429, f334])).
fof(f334, plain, ! [X0] : (r1(X0, sK58(X0)) | ~ p504(X0) | ~ sP22(X0)), inference(cnf_transformation, [], [f123])).
fof(f2312, plain, (~ spl92_3 | ~ spl92_23), inference(avatar_contradiction_clause, [], [f2311])).
fof(f2311, plain, ($false | (~ spl92_3 | ~ spl92_23)), inference(subsumption_resolution, [], [f2310, f555])).
fof(f555, plain, sP32(sK91), inference(resolution, [], [f523, f266])).
fof(f266, plain, ! [X0] : (~ sP40(X0) | sP32(X0)), inference(cnf_transformation, [], [f51])).
fof(f2310, plain, (~ sP32(sK91) | (~ spl92_3 | ~ spl92_23)), inference(subsumption_resolution, [], [f2309, f492])).
fof(f2309, plain, (~ p503(sK91) | ~ sP32(sK91) | (~ spl92_3 | ~ spl92_23)), inference(resolution, [], [f1506, f315])).
fof(f315, plain, ! [X0] : (~ p103(sK48(X0)) | ~ p503(X0) | ~ sP32(X0)), inference(cnf_transformation, [], [f83])).
fof(f83, plain, ! [X0] : ((~ p103(sK48(X0)) & r1(X0, sK48(X0))) | ~ p503(X0) | ~ sP32(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK48])], [f81, f82])).
fof(f82, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) => (~ p103(sK48(X0)) & r1(X0, sK48(X0)))), introduced(choice_axiom, [])).
fof(f81, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) | ~ p503(X0) | ~ sP32(X0)), inference(rectify, [], [f80])).
fof(f80, plain, ! [X11] : (? [X21] : (~ p103(X21) & r1(X11, X21)) | ~ p503(X11) | ~ sP32(X11)), inference(nnf_transformation, [], [f40])).
fof(f1506, plain, (p103(sK48(sK91)) | (~ spl92_3 | ~ spl92_23)), inference(subsumption_resolution, [], [f1505, f555])).
fof(f1505, plain, (p103(sK48(sK91)) | ~ sP32(sK91) | (~ spl92_3 | ~ spl92_23)), inference(subsumption_resolution, [], [f1482, f492])).
fof(f1482, plain, (p103(sK48(sK91)) | ~ p503(sK91) | ~ sP32(sK91) | ~ spl92_3), inference(resolution, [], [f415, f314])).
fof(f314, plain, ! [X0] : (r1(X0, sK48(X0)) | ~ p503(X0) | ~ sP32(X0)), inference(cnf_transformation, [], [f83])).
fof(f2290, plain, (spl92_184 | ~ spl92_185 | ~ spl92_3), inference(avatar_split_clause, [], [f2281, f414, f2287, f2283])).
fof(f2281, plain, (~ p203(sK72(sK91)) | p103(sK71(sK91)) | ~ spl92_3), inference(subsumption_resolution, [], [f2277, f558])).
fof(f2277, plain, (~ p203(sK72(sK91)) | ~ sP9(sK91) | p103(sK71(sK91)) | ~ spl92_3), inference(resolution, [], [f361, f415])).
fof(f361, plain, ! [X0] : (r1(X0, sK71(X0)) | ~ p203(sK72(X0)) | ~ sP9(X0)), inference(cnf_transformation, [], [f176])).
fof(f1473, plain, (~ spl92_1 | ~ spl92_26), inference(avatar_contradiction_clause, [], [f1472])).
fof(f1472, plain, ($false | (~ spl92_1 | ~ spl92_26)), inference(subsumption_resolution, [], [f1471, f533])).
fof(f533, plain, sP16(sK91), inference(resolution, [], [f523, f235])).
fof(f235, plain, ! [X0] : (~ sP40(X0) | sP16(X0)), inference(cnf_transformation, [], [f51])).
fof(f1471, plain, (~ sP16(sK91) | (~ spl92_1 | ~ spl92_26)), inference(subsumption_resolution, [], [f1470, f505])).
fof(f1470, plain, (~ p605(sK91) | ~ sP16(sK91) | (~ spl92_1 | ~ spl92_26)), inference(resolution, [], [f1419, f347])).
fof(f347, plain, ! [X0] : (~ p105(sK64(X0)) | ~ p605(X0) | ~ sP16(X0)), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ! [X0] : ((~ p105(sK64(X0)) & r1(X0, sK64(X0))) | ~ p605(X0) | ~ sP16(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK64])], [f145, f146])).
fof(f146, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) => (~ p105(sK64(X0)) & r1(X0, sK64(X0)))), introduced(choice_axiom, [])).
fof(f145, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) | ~ p605(X0) | ~ sP16(X0)), inference(rectify, [], [f144])).
fof(f144, plain, ! [X11] : (? [X49] : (~ p105(X49) & r1(X11, X49)) | ~ p605(X11) | ~ sP16(X11)), inference(nnf_transformation, [], [f24])).
fof(f1419, plain, (p105(sK64(sK91)) | (~ spl92_1 | ~ spl92_26)), inference(subsumption_resolution, [], [f1418, f533])).
fof(f1418, plain, (~ sP16(sK91) | p105(sK64(sK91)) | (~ spl92_1 | ~ spl92_26)), inference(subsumption_resolution, [], [f1412, f505])).
fof(f1412, plain, (~ p605(sK91) | ~ sP16(sK91) | p105(sK64(sK91)) | ~ spl92_1), inference(resolution, [], [f346, f409])).
fof(f346, plain, ! [X0] : (r1(X0, sK64(X0)) | ~ p605(X0) | ~ sP16(X0)), inference(cnf_transformation, [], [f147])).
fof(f801, plain, (~ spl92_2 | ~ spl92_22), inference(avatar_contradiction_clause, [], [f800])).
fof(f800, plain, ($false | (~ spl92_2 | ~ spl92_22)), inference(subsumption_resolution, [], [f799, f546])).
fof(f546, plain, sP25(sK91), inference(resolution, [], [f523, f251])).
fof(f251, plain, ! [X0] : (~ sP40(X0) | sP25(X0)), inference(cnf_transformation, [], [f51])).
fof(f799, plain, (~ sP25(sK91) | (~ spl92_2 | ~ spl92_22)), inference(subsumption_resolution, [], [f798, f488])).
fof(f488, plain, (p504(sK91) | ~ spl92_22), inference(avatar_component_clause, [], [f486])).
fof(f798, plain, (~ p504(sK91) | ~ sP25(sK91) | (~ spl92_2 | ~ spl92_22)), inference(resolution, [], [f765, f329])).
fof(f329, plain, ! [X0] : (~ p104(sK55(X0)) | ~ p504(X0) | ~ sP25(X0)), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ! [X0] : ((~ p104(sK55(X0)) & r1(X0, sK55(X0))) | ~ p504(X0) | ~ sP25(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK55])], [f109, f110])).
fof(f110, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) => (~ p104(sK55(X0)) & r1(X0, sK55(X0)))), introduced(choice_axiom, [])).
fof(f109, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) | ~ p504(X0) | ~ sP25(X0)), inference(rectify, [], [f108])).
fof(f108, plain, ! [X11] : (? [X32] : (~ p104(X32) & r1(X11, X32)) | ~ p504(X11) | ~ sP25(X11)), inference(nnf_transformation, [], [f33])).
fof(f765, plain, (p104(sK55(sK91)) | (~ spl92_2 | ~ spl92_22)), inference(subsumption_resolution, [], [f764, f546])).
fof(f764, plain, (~ sP25(sK91) | p104(sK55(sK91)) | (~ spl92_2 | ~ spl92_22)), inference(subsumption_resolution, [], [f756, f488])).
fof(f756, plain, (~ p504(sK91) | ~ sP25(sK91) | p104(sK55(sK91)) | ~ spl92_2), inference(resolution, [], [f328, f412])).
fof(f328, plain, ! [X0] : (r1(X0, sK55(X0)) | ~ p504(X0) | ~ sP25(X0)), inference(cnf_transformation, [], [f111])).
fof(f618, plain, (~ spl92_5 | ~ spl92_30), inference(avatar_split_clause, [], [f617, f519, f420])).
fof(f519, plain, (spl92_30 <=> p601(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_30])])).
fof(f617, plain, (~ p601(sK91) | ~ p101(sK91)), inference(resolution, [], [f295, f523])).
fof(f295, plain, ! [X0] : (~ sP40(X0) | ~ p601(X0) | ~ p101(X0)), inference(cnf_transformation, [], [f51])).
fof(f613, plain, (~ spl92_10 | ~ spl92_30), inference(avatar_split_clause, [], [f612, f519, f438])).
fof(f612, plain, (~ p601(sK91) | ~ p201(sK91)), inference(resolution, [], [f291, f523])).
fof(f291, plain, ! [X0] : (~ sP40(X0) | ~ p601(X0) | ~ p201(X0)), inference(cnf_transformation, [], [f51])).
fof(f609, plain, (~ spl92_15 | ~ spl92_30), inference(avatar_split_clause, [], [f608, f519, f457])).
fof(f608, plain, (~ p601(sK91) | ~ p301(sK91)), inference(resolution, [], [f288, f523])).
fof(f288, plain, ! [X0] : (~ sP40(X0) | ~ p601(X0) | ~ p301(X0)), inference(cnf_transformation, [], [f51])).
fof(f606, plain, (~ spl92_20 | ~ spl92_30), inference(avatar_split_clause, [], [f605, f519, f477])).
fof(f605, plain, (~ p601(sK91) | ~ p401(sK91)), inference(resolution, [], [f286, f523])).
fof(f286, plain, ! [X0] : (~ sP40(X0) | ~ p601(X0) | ~ p401(X0)), inference(cnf_transformation, [], [f51])).
fof(f604, plain, (~ spl92_25 | ~ spl92_30), inference(avatar_split_clause, [], [f603, f519, f498])).
fof(f603, plain, (~ p601(sK91) | ~ p501(sK91)), inference(resolution, [], [f285, f523])).
fof(f285, plain, ! [X0] : (~ sP40(X0) | ~ p601(X0) | ~ p501(X0)), inference(cnf_transformation, [], [f51])).
fof(f599, plain, (~ spl92_9 | ~ spl92_29), inference(avatar_split_clause, [], [f598, f515, f434])).
fof(f598, plain, (~ p602(sK91) | ~ p202(sK91)), inference(resolution, [], [f276, f523])).
fof(f276, plain, ! [X0] : (~ sP40(X0) | ~ p602(X0) | ~ p202(X0)), inference(cnf_transformation, [], [f51])).
fof(f595, plain, (~ spl92_14 | ~ spl92_29), inference(avatar_split_clause, [], [f594, f515, f453])).
fof(f594, plain, (~ p602(sK91) | ~ p302(sK91)), inference(resolution, [], [f273, f523])).
fof(f273, plain, ! [X0] : (~ sP40(X0) | ~ p602(X0) | ~ p302(X0)), inference(cnf_transformation, [], [f51])).
fof(f592, plain, (~ spl92_19 | ~ spl92_29), inference(avatar_split_clause, [], [f591, f515, f473])).
fof(f591, plain, (~ p602(sK91) | ~ p402(sK91)), inference(resolution, [], [f271, f523])).
fof(f271, plain, ! [X0] : (~ sP40(X0) | ~ p602(X0) | ~ p402(X0)), inference(cnf_transformation, [], [f51])).
fof(f590, plain, (~ spl92_24 | ~ spl92_29), inference(avatar_split_clause, [], [f589, f515, f494])).
fof(f589, plain, (~ p602(sK91) | ~ p502(sK91)), inference(resolution, [], [f270, f523])).
fof(f270, plain, ! [X0] : (~ sP40(X0) | ~ p602(X0) | ~ p502(X0)), inference(cnf_transformation, [], [f51])).
fof(f586, plain, (~ spl92_13 | ~ spl92_28), inference(avatar_split_clause, [], [f585, f511, f449])).
fof(f585, plain, (~ p603(sK91) | ~ p303(sK91)), inference(resolution, [], [f258, f523])).
fof(f258, plain, ! [X0] : (~ sP40(X0) | ~ p603(X0) | ~ p303(X0)), inference(cnf_transformation, [], [f51])).
fof(f583, plain, (~ spl92_18 | ~ spl92_28), inference(avatar_split_clause, [], [f582, f511, f469])).
fof(f582, plain, (~ p603(sK91) | ~ p403(sK91)), inference(resolution, [], [f256, f523])).
fof(f256, plain, ! [X0] : (~ sP40(X0) | ~ p603(X0) | ~ p403(X0)), inference(cnf_transformation, [], [f51])).
fof(f581, plain, (~ spl92_23 | ~ spl92_28), inference(avatar_split_clause, [], [f580, f511, f490])).
fof(f580, plain, (~ p603(sK91) | ~ p503(sK91)), inference(resolution, [], [f255, f523])).
fof(f255, plain, ! [X0] : (~ sP40(X0) | ~ p603(X0) | ~ p503(X0)), inference(cnf_transformation, [], [f51])).
fof(f575, plain, (~ spl92_27 | ~ spl92_22), inference(avatar_split_clause, [], [f574, f486, f507])).
fof(f574, plain, (~ p604(sK91) | ~ spl92_22), inference(subsumption_resolution, [], [f573, f488])).
fof(f573, plain, (~ p604(sK91) | ~ p504(sK91)), inference(resolution, [], [f240, f523])).
fof(f240, plain, ! [X0] : (~ sP40(X0) | ~ p604(X0) | ~ p504(X0)), inference(cnf_transformation, [], [f51])).
fof(f522, plain, (spl92_26 | spl92_27 | spl92_28 | spl92_29 | spl92_30), inference(avatar_split_clause, [], [f401, f519, f515, f511, f507, f503])).
fof(f401, plain, (p601(sK91) | p602(sK91) | p603(sK91) | p604(sK91) | p605(sK91)), inference(cnf_transformation, [], [f223])).
fof(f501, plain, (spl92_21 | spl92_22 | spl92_23 | spl92_24 | spl92_25), inference(avatar_split_clause, [], [f402, f498, f494, f490, f486, f482])).
fof(f402, plain, (p501(sK91) | p502(sK91) | p503(sK91) | p504(sK91) | p505(sK91)), inference(cnf_transformation, [], [f223])).
fof(f480, plain, (spl92_16 | spl92_17 | spl92_18 | spl92_19 | spl92_20), inference(avatar_split_clause, [], [f403, f477, f473, f469, f465, f462])).
fof(f403, plain, ! [X10] : (p401(sK91) | p402(sK91) | p403(sK91) | p404(sK91) | p405(X10) | ~ r1(sK91, X10)), inference(cnf_transformation, [], [f223])).
fof(f460, plain, (spl92_11 | spl92_12 | spl92_13 | spl92_14 | spl92_15), inference(avatar_split_clause, [], [f404, f457, f453, f449, f446, f443])).
fof(f404, plain, ! [X8, X9] : (p301(sK91) | p302(sK91) | p303(sK91) | p304(X8) | ~ r1(sK91, X8) | p305(X9) | ~ r1(sK91, X9)), inference(cnf_transformation, [], [f223])).
fof(f441, plain, (spl92_6 | spl92_7 | spl92_8 | spl92_9 | spl92_10), inference(avatar_split_clause, [], [f405, f438, f434, f431, f428, f425])).
fof(f405, plain, ! [X6, X7, X5] : (p201(sK91) | p202(sK91) | p203(X5) | ~ r1(sK91, X5) | p204(X6) | ~ r1(sK91, X6) | p205(X7) | ~ r1(sK91, X7)), inference(cnf_transformation, [], [f223])).
fof(f423, plain, (spl92_1 | spl92_2 | spl92_3 | spl92_4 | spl92_5), inference(avatar_split_clause, [], [f406, f420, f417, f414, f411, f408])).
fof(f406, plain, ! [X4, X2, X3, X1] : (p101(sK91) | p102(X1) | ~ r1(sK91, X1) | p103(X2) | ~ r1(sK91, X2) | p104(X3) | ~ r1(sK91, X3) | p105(X4) | ~ r1(sK91, X4)), inference(cnf_transformation, [], [f223])).