fof(f144, plain, $false, inference(subsumption_resolution, [], [f143, f36])).
fof(f36, plain, sP0(sK3), inference(resolution, [], [f27, f19])).
fof(f19, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL656+1.001.p', reflexivity)).
fof(f27, plain, ! [X1] : (~ r1(sK3, X1) | sP0(X1)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, (p100(sK3) & ~ p101(sK3) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & sP0(X1)) | ~ r1(sK3, X1)) & ! [X6] : (p2(X6) | ~ r1(sK3, X6))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3])], [f16, f17])).
fof(f17, plain, (? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & sP0(X1)) | ~ r1(X0, X1)) & ! [X6] : (p2(X6) | ~ r1(X0, X6))) => (p100(sK3) & ~ p101(sK3) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & sP0(X1)) | ~ r1(sK3, X1)) & ! [X6] : (p2(X6) | ~ r1(sK3, X6)))), introduced(choice_axiom, [])).
fof(f16, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & sP0(X1)) | ~ r1(X0, X1)) & ! [X6] : (p2(X6) | ~ r1(X0, X6))), inference(rectify, [], [f10])).
fof(f10, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & sP0(X1)) | ~ r1(X0, X1)) & ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(definition_folding, [], [f8, e9])).
fof(f9, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X6] : (p101(X6) & p2(X6) & r1(X1, X6)) & ? [X7] : (p101(X7) & ~ p2(X7) & r1(X1, X7))) | ~ sP0(X1)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X1] : (sP0(X1) <=> (~ p100(X1) | p101(X1) | (? [X6] : (p101(X6) & p2(X6) & r1(X1, X6)) & ? [X7] : (p101(X7) & ~ p2(X7) & r1(X1, X7))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f8, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p100(X1) | p101(X1) | (? [X6] : (p101(X6) & p2(X6) & r1(X1, X6)) & ? [X7] : (p101(X7) & ~ p2(X7) & r1(X1, X7))))) | ~ r1(X0, X1)) & ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(flattening, [], [f7])).
fof(f7, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & ((~ p100(X1) | p101(X1)) | (? [X6] : ((p101(X6) & p2(X6)) & r1(X1, X6)) & ? [X7] : ((p101(X7) & ~ p2(X7)) & r1(X1, X7))))) | ~ r1(X0, X1))) & ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X6] : (~ (p101(X6) & p2(X6)) | ~ r1(X1, X6)) & ~ ! [X7] : (~ (p101(X7) & ~ p2(X7)) | ~ r1(X1, X7))))) | ~ r1(X0, X1))) | ~ ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(pure_predicate_removal, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X6] : (~ (p101(X6) & ~ p102(X6) & p2(X6)) | ~ r1(X1, X6)) & ~ ! [X7] : (~ (p101(X7) & ~ p102(X7) & ~ p2(X7)) | ~ r1(X1, X7))))) | ~ r1(X0, X1))) | ~ ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(flattening, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X6] : (~ (p101(X6) & ~ p102(X6) & p2(X6)) | ~ r1(X1, X6)) & ~ ! [X7] : (~ (p101(X7) & ~ p102(X7) & ~ p2(X7)) | ~ r1(X1, X7))))) | ~ r1(X0, X1))) | ~ ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(rectify, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p2(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p2(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL656+1.001.p', main)).
fof(f143, plain, ~ sP0(sK3), inference(subsumption_resolution, [], [f142, f34])).
fof(f34, plain, p100(sK3), inference(cnf_transformation, [], [f18])).
fof(f142, plain, (~ p100(sK3) | ~ sP0(sK3)), inference(subsumption_resolution, [], [f141, f33])).
fof(f33, plain, ~ p101(sK3), inference(cnf_transformation, [], [f18])).
fof(f141, plain, (p101(sK3) | ~ p100(sK3) | ~ sP0(sK3)), inference(resolution, [], [f58, f21])).
fof(f21, plain, ! [X0] : (~ p2(sK2(X0)) | p101(X0) | ~ p100(X0) | ~ sP0(X0)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ! [X0] : (~ p100(X0) | p101(X0) | ((p101(sK1(X0)) & p2(sK1(X0)) & r1(X0, sK1(X0))) & (p101(sK2(X0)) & ~ p2(sK2(X0)) & r1(X0, sK2(X0)))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2])], [f12, f14, f13])).
fof(f13, plain, ! [X0] : (? [X1] : (p101(X1) & p2(X1) & r1(X0, X1)) => (p101(sK1(X0)) & p2(sK1(X0)) & r1(X0, sK1(X0)))), introduced(choice_axiom, [])).
fof(f14, plain, ! [X0] : (? [X2] : (p101(X2) & ~ p2(X2) & r1(X0, X2)) => (p101(sK2(X0)) & ~ p2(sK2(X0)) & r1(X0, sK2(X0)))), introduced(choice_axiom, [])).
fof(f12, plain, ! [X0] : (~ p100(X0) | p101(X0) | (? [X1] : (p101(X1) & p2(X1) & r1(X0, X1)) & ? [X2] : (p101(X2) & ~ p2(X2) & r1(X0, X2))) | ~ sP0(X0)), inference(rectify, [], [f11])).
fof(f11, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X6] : (p101(X6) & p2(X6) & r1(X1, X6)) & ? [X7] : (p101(X7) & ~ p2(X7) & r1(X1, X7))) | ~ sP0(X1)), inference(nnf_transformation, [], [f9])).
fof(f58, plain, p2(sK2(sK3)), inference(subsumption_resolution, [], [f57, f36])).
fof(f57, plain, (~ sP0(sK3) | p2(sK2(sK3))), inference(subsumption_resolution, [], [f56, f34])).
fof(f56, plain, (~ p100(sK3) | ~ sP0(sK3) | p2(sK2(sK3))), inference(subsumption_resolution, [], [f40, f33])).
fof(f40, plain, (p101(sK3) | ~ p100(sK3) | ~ sP0(sK3) | p2(sK2(sK3))), inference(resolution, [], [f20, f26])).
fof(f26, plain, ! [X6] : (~ r1(sK3, X6) | p2(X6)), inference(cnf_transformation, [], [f18])).
fof(f20, plain, ! [X0] : (r1(X0, sK2(X0)) | p101(X0) | ~ p100(X0) | ~ sP0(X0)), inference(cnf_transformation, [], [f15])).