fof(f13897, plain, $false, inference(avatar_sat_refutation, [], [f3064, f3983, f13896])).
fof(f13896, plain, (spl63_44 | ~ spl63_45), inference(avatar_contradiction_clause, [], [f13895])).
fof(f13895, plain, ($false | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f13894, f157])).
fof(f157, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.015.p', reflexivity)).
fof(f13894, plain, (~ r1(sK62, sK62) | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f13893, f357])).
fof(f357, plain, ~ p101(sK62), inference(cnf_transformation, [], [f156])).
fof(f156, plain, (p100(sK62) & ~ p101(sK62) & ! [X1] : (sP31(X1) | ~ r1(sK62, X1)) & ! [X2] : (p7(X2) | ~ r1(sK62, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK62])], [f154, f155])).
fof(f155, plain, (? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP31(X1) | ~ r1(X0, X1)) & ! [X2] : (p7(X2) | ~ r1(X0, X2))) => (p100(sK62) & ~ p101(sK62) & ! [X1] : (sP31(X1) | ~ r1(sK62, X1)) & ! [X2] : (p7(X2) | ~ r1(sK62, X2)))), introduced(choice_axiom, [])).
fof(f154, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP31(X1) | ~ r1(X0, X1)) & ! [X2] : (p7(X2) | ~ r1(X0, X2))), inference(rectify, [], [f44])).
fof(f44, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP31(X1) | ~ r1(X0, X1)) & ! [X64] : (p7(X64) | ~ r1(X0, X64))), inference(definition_folding, [], [f11, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12])).
fof(f12, plain, ! [X1] : (~ p113(X1) | p114(X1) | (? [X60] : (p114(X60) & ~ p115(X60) & p15(X60) & r1(X1, X60)) & ? [X61] : (p114(X61) & ~ p115(X61) & ~ p15(X61) & r1(X1, X61))) | ~ sP0(X1)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X1] : (sP0(X1) <=> (~ p113(X1) | p114(X1) | (? [X60] : (p114(X60) & ~ p115(X60) & p15(X60) & r1(X1, X60)) & ? [X61] : (p114(X61) & ~ p115(X61) & ~ p15(X61) & r1(X1, X61))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f13, plain, ! [X1] : (~ p112(X1) | p113(X1) | (? [X58] : (p113(X58) & ~ p114(X58) & p14(X58) & r1(X1, X58)) & ? [X59] : (p113(X59) & ~ p114(X59) & ~ p14(X59) & r1(X1, X59))) | ~ sP1(X1)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X1] : (sP1(X1) <=> (~ p112(X1) | p113(X1) | (? [X58] : (p113(X58) & ~ p114(X58) & p14(X58) & r1(X1, X58)) & ? [X59] : (p113(X59) & ~ p114(X59) & ~ p14(X59) & r1(X1, X59))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f14, plain, ! [X1] : (~ p111(X1) | p112(X1) | (? [X56] : (p112(X56) & ~ p113(X56) & p13(X56) & r1(X1, X56)) & ? [X57] : (p112(X57) & ~ p113(X57) & ~ p13(X57) & r1(X1, X57))) | ~ sP2(X1)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X1] : (sP2(X1) <=> (~ p111(X1) | p112(X1) | (? [X56] : (p112(X56) & ~ p113(X56) & p13(X56) & r1(X1, X56)) & ? [X57] : (p112(X57) & ~ p113(X57) & ~ p13(X57) & r1(X1, X57))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f15, plain, ! [X1] : (~ p110(X1) | p111(X1) | (? [X54] : (p111(X54) & ~ p112(X54) & p12(X54) & r1(X1, X54)) & ? [X55] : (p111(X55) & ~ p112(X55) & ~ p12(X55) & r1(X1, X55))) | ~ sP3(X1)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X1] : (sP3(X1) <=> (~ p110(X1) | p111(X1) | (? [X54] : (p111(X54) & ~ p112(X54) & p12(X54) & r1(X1, X54)) & ? [X55] : (p111(X55) & ~ p112(X55) & ~ p12(X55) & r1(X1, X55))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f16, plain, ! [X1] : (~ p109(X1) | p110(X1) | (? [X52] : (p110(X52) & ~ p111(X52) & p11(X52) & r1(X1, X52)) & ? [X53] : (p110(X53) & ~ p111(X53) & ~ p11(X53) & r1(X1, X53))) | ~ sP4(X1)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X1] : (sP4(X1) <=> (~ p109(X1) | p110(X1) | (? [X52] : (p110(X52) & ~ p111(X52) & p11(X52) & r1(X1, X52)) & ? [X53] : (p110(X53) & ~ p111(X53) & ~ p11(X53) & r1(X1, X53))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f17, plain, ! [X1] : (~ p108(X1) | p109(X1) | (? [X50] : (p109(X50) & ~ p110(X50) & p10(X50) & r1(X1, X50)) & ? [X51] : (p109(X51) & ~ p110(X51) & ~ p10(X51) & r1(X1, X51))) | ~ sP5(X1)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X1] : (sP5(X1) <=> (~ p108(X1) | p109(X1) | (? [X50] : (p109(X50) & ~ p110(X50) & p10(X50) & r1(X1, X50)) & ? [X51] : (p109(X51) & ~ p110(X51) & ~ p10(X51) & r1(X1, X51))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f18, plain, ! [X1] : (~ p107(X1) | p108(X1) | (? [X48] : (p108(X48) & ~ p109(X48) & p9(X48) & r1(X1, X48)) & ? [X49] : (p108(X49) & ~ p109(X49) & ~ p9(X49) & r1(X1, X49))) | ~ sP6(X1)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X1] : (sP6(X1) <=> (~ p107(X1) | p108(X1) | (? [X48] : (p108(X48) & ~ p109(X48) & p9(X48) & r1(X1, X48)) & ? [X49] : (p108(X49) & ~ p109(X49) & ~ p9(X49) & r1(X1, X49))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f19, plain, ! [X1] : (~ p106(X1) | p107(X1) | (? [X46] : (p107(X46) & ~ p108(X46) & p8(X46) & r1(X1, X46)) & ? [X47] : (p107(X47) & ~ p108(X47) & ~ p8(X47) & r1(X1, X47))) | ~ sP7(X1)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X1] : (sP7(X1) <=> (~ p106(X1) | p107(X1) | (? [X46] : (p107(X46) & ~ p108(X46) & p8(X46) & r1(X1, X46)) & ? [X47] : (p107(X47) & ~ p108(X47) & ~ p8(X47) & r1(X1, X47))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f20, plain, ! [X1] : (~ p105(X1) | p106(X1) | (? [X44] : (p106(X44) & ~ p107(X44) & p7(X44) & r1(X1, X44)) & ? [X45] : (p106(X45) & ~ p107(X45) & ~ p7(X45) & r1(X1, X45))) | ~ sP8(X1)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X1] : (sP8(X1) <=> (~ p105(X1) | p106(X1) | (? [X44] : (p106(X44) & ~ p107(X44) & p7(X44) & r1(X1, X44)) & ? [X45] : (p106(X45) & ~ p107(X45) & ~ p7(X45) & r1(X1, X45))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f21, plain, ! [X1] : (~ p104(X1) | p105(X1) | (? [X42] : (p105(X42) & ~ p106(X42) & p6(X42) & r1(X1, X42)) & ? [X43] : (p105(X43) & ~ p106(X43) & ~ p6(X43) & r1(X1, X43))) | ~ sP9(X1)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X1] : (sP9(X1) <=> (~ p104(X1) | p105(X1) | (? [X42] : (p105(X42) & ~ p106(X42) & p6(X42) & r1(X1, X42)) & ? [X43] : (p105(X43) & ~ p106(X43) & ~ p6(X43) & r1(X1, X43))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f22, plain, ! [X1] : (~ p103(X1) | p104(X1) | (? [X40] : (p104(X40) & ~ p105(X40) & p5(X40) & r1(X1, X40)) & ? [X41] : (p104(X41) & ~ p105(X41) & ~ p5(X41) & r1(X1, X41))) | ~ sP10(X1)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X1] : (sP10(X1) <=> (~ p103(X1) | p104(X1) | (? [X40] : (p104(X40) & ~ p105(X40) & p5(X40) & r1(X1, X40)) & ? [X41] : (p104(X41) & ~ p105(X41) & ~ p5(X41) & r1(X1, X41))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f23, plain, ! [X1] : (~ p102(X1) | p103(X1) | (? [X38] : (p103(X38) & ~ p104(X38) & p4(X38) & r1(X1, X38)) & ? [X39] : (p103(X39) & ~ p104(X39) & ~ p4(X39) & r1(X1, X39))) | ~ sP11(X1)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X1] : (sP11(X1) <=> (~ p102(X1) | p103(X1) | (? [X38] : (p103(X38) & ~ p104(X38) & p4(X38) & r1(X1, X38)) & ? [X39] : (p103(X39) & ~ p104(X39) & ~ p4(X39) & r1(X1, X39))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f24, plain, ! [X1] : (~ p101(X1) | p102(X1) | (? [X36] : (p102(X36) & ~ p103(X36) & p3(X36) & r1(X1, X36)) & ? [X37] : (p102(X37) & ~ p103(X37) & ~ p3(X37) & r1(X1, X37))) | ~ sP12(X1)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X1] : (sP12(X1) <=> (~ p101(X1) | p102(X1) | (? [X36] : (p102(X36) & ~ p103(X36) & p3(X36) & r1(X1, X36)) & ? [X37] : (p102(X37) & ~ p103(X37) & ~ p3(X37) & r1(X1, X37))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f25, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X34] : (p101(X34) & ~ p102(X34) & p2(X34) & r1(X1, X34)) & ? [X35] : (p101(X35) & ~ p102(X35) & ~ p2(X35) & r1(X1, X35))) | ~ sP13(X1)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X1] : (sP13(X1) <=> (~ p100(X1) | p101(X1) | (? [X34] : (p101(X34) & ~ p102(X34) & p2(X34) & r1(X1, X34)) & ? [X35] : (p101(X35) & ~ p102(X35) & ~ p2(X35) & r1(X1, X35))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f26, plain, ! [X1] : (~ p114(X1) | p115(X1) | (? [X62] : (p115(X62) & p16(X62) & r1(X1, X62)) & ? [X63] : (p115(X63) & ~ p16(X63) & r1(X1, X63))) | ~ sP14(X1)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X1] : (sP14(X1) <=> (~ p114(X1) | p115(X1) | (? [X62] : (p115(X62) & p16(X62) & r1(X1, X62)) & ? [X63] : (p115(X63) & ~ p16(X63) & r1(X1, X63))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f27, plain, ! [X1] : (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33)))) | ~ sP15(X1)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X1] : (sP15(X1) <=> (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f28, plain, ! [X1] : (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31)))) | ~ sP16(X1)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X1] : (sP16(X1) <=> (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f29, plain, ! [X1] : (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29)))) | ~ sP17(X1)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X1] : (sP17(X1) <=> (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f30, plain, ! [X1] : (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27)))) | ~ sP18(X1)), inference(usedef, [], [e30])).
fof(e30, plain, ! [X1] : (sP18(X1) <=> (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f31, plain, ! [X1] : (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25)))) | ~ sP19(X1)), inference(usedef, [], [e31])).
fof(e31, plain, ! [X1] : (sP19(X1) <=> (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f32, plain, ! [X1] : (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23)))) | ~ sP20(X1)), inference(usedef, [], [e32])).
fof(e32, plain, ! [X1] : (sP20(X1) <=> (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f33, plain, ! [X1] : (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21)))) | ~ sP21(X1)), inference(usedef, [], [e33])).
fof(e33, plain, ! [X1] : (sP21(X1) <=> (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f34, plain, ! [X1] : (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19)))) | ~ sP22(X1)), inference(usedef, [], [e34])).
fof(e34, plain, ! [X1] : (sP22(X1) <=> (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f35, plain, ! [X1] : (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17)))) | ~ sP23(X1)), inference(usedef, [], [e35])).
fof(e35, plain, ! [X1] : (sP23(X1) <=> (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f36, plain, ! [X1] : (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15)))) | ~ sP24(X1)), inference(usedef, [], [e36])).
fof(e36, plain, ! [X1] : (sP24(X1) <=> (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f37, plain, ! [X1] : (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13)))) | ~ sP25(X1)), inference(usedef, [], [e37])).
fof(e37, plain, ! [X1] : (sP25(X1) <=> (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f38, plain, ! [X1] : (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11)))) | ~ sP26(X1)), inference(usedef, [], [e38])).
fof(e38, plain, ! [X1] : (sP26(X1) <=> (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f39, plain, ! [X1] : (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9)))) | ~ sP27(X1)), inference(usedef, [], [e39])).
fof(e39, plain, ! [X1] : (sP27(X1) <=> (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f40, plain, ! [X1] : (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7)))) | ~ sP28(X1)), inference(usedef, [], [e40])).
fof(e40, plain, ! [X1] : (sP28(X1) <=> (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f41, plain, ! [X1] : (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5)))) | ~ sP29(X1)), inference(usedef, [], [e41])).
fof(e41, plain, ! [X1] : (sP29(X1) <=> (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f42, plain, ! [X1] : (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3)))) | ~ sP30(X1)), inference(usedef, [], [e42])).
fof(e42, plain, ! [X1] : (sP30(X1) <=> (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f43, plain, ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & sP30(X1) & sP29(X1) & sP28(X1) & sP27(X1) & sP26(X1) & sP25(X1) & sP24(X1) & sP23(X1) & sP22(X1) & sP21(X1) & sP20(X1) & sP19(X1) & sP18(X1) & sP17(X1) & sP16(X1) & sP15(X1) & sP13(X1) & sP12(X1) & sP11(X1) & sP10(X1) & sP9(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP4(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP14(X1)) | ~ sP31(X1)), inference(usedef, [], [e43])).
fof(e43, plain, ! [X1] : (sP31(X1) <=> ((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & sP30(X1) & sP29(X1) & sP28(X1) & sP27(X1) & sP26(X1) & sP25(X1) & sP24(X1) & sP23(X1) & sP22(X1) & sP21(X1) & sP20(X1) & sP19(X1) & sP18(X1) & sP17(X1) & sP16(X1) & sP15(X1) & sP13(X1) & sP12(X1) & sP11(X1) & sP10(X1) & sP9(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP4(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP14(X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f11, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25))))) & (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27))))) & (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29))))) & (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31))))) & (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33))))) & (~ p100(X1) | p101(X1) | (? [X34] : (p101(X34) & ~ p102(X34) & p2(X34) & r1(X1, X34)) & ? [X35] : (p101(X35) & ~ p102(X35) & ~ p2(X35) & r1(X1, X35)))) & (~ p101(X1) | p102(X1) | (? [X36] : (p102(X36) & ~ p103(X36) & p3(X36) & r1(X1, X36)) & ? [X37] : (p102(X37) & ~ p103(X37) & ~ p3(X37) & r1(X1, X37)))) & (~ p102(X1) | p103(X1) | (? [X38] : (p103(X38) & ~ p104(X38) & p4(X38) & r1(X1, X38)) & ? [X39] : (p103(X39) & ~ p104(X39) & ~ p4(X39) & r1(X1, X39)))) & (~ p103(X1) | p104(X1) | (? [X40] : (p104(X40) & ~ p105(X40) & p5(X40) & r1(X1, X40)) & ? [X41] : (p104(X41) & ~ p105(X41) & ~ p5(X41) & r1(X1, X41)))) & (~ p104(X1) | p105(X1) | (? [X42] : (p105(X42) & ~ p106(X42) & p6(X42) & r1(X1, X42)) & ? [X43] : (p105(X43) & ~ p106(X43) & ~ p6(X43) & r1(X1, X43)))) & (~ p105(X1) | p106(X1) | (? [X44] : (p106(X44) & ~ p107(X44) & p7(X44) & r1(X1, X44)) & ? [X45] : (p106(X45) & ~ p107(X45) & ~ p7(X45) & r1(X1, X45)))) & (~ p106(X1) | p107(X1) | (? [X46] : (p107(X46) & ~ p108(X46) & p8(X46) & r1(X1, X46)) & ? [X47] : (p107(X47) & ~ p108(X47) & ~ p8(X47) & r1(X1, X47)))) & (~ p107(X1) | p108(X1) | (? [X48] : (p108(X48) & ~ p109(X48) & p9(X48) & r1(X1, X48)) & ? [X49] : (p108(X49) & ~ p109(X49) & ~ p9(X49) & r1(X1, X49)))) & (~ p108(X1) | p109(X1) | (? [X50] : (p109(X50) & ~ p110(X50) & p10(X50) & r1(X1, X50)) & ? [X51] : (p109(X51) & ~ p110(X51) & ~ p10(X51) & r1(X1, X51)))) & (~ p109(X1) | p110(X1) | (? [X52] : (p110(X52) & ~ p111(X52) & p11(X52) & r1(X1, X52)) & ? [X53] : (p110(X53) & ~ p111(X53) & ~ p11(X53) & r1(X1, X53)))) & (~ p110(X1) | p111(X1) | (? [X54] : (p111(X54) & ~ p112(X54) & p12(X54) & r1(X1, X54)) & ? [X55] : (p111(X55) & ~ p112(X55) & ~ p12(X55) & r1(X1, X55)))) & (~ p111(X1) | p112(X1) | (? [X56] : (p112(X56) & ~ p113(X56) & p13(X56) & r1(X1, X56)) & ? [X57] : (p112(X57) & ~ p113(X57) & ~ p13(X57) & r1(X1, X57)))) & (~ p112(X1) | p113(X1) | (? [X58] : (p113(X58) & ~ p114(X58) & p14(X58) & r1(X1, X58)) & ? [X59] : (p113(X59) & ~ p114(X59) & ~ p14(X59) & r1(X1, X59)))) & (~ p113(X1) | p114(X1) | (? [X60] : (p114(X60) & ~ p115(X60) & p15(X60) & r1(X1, X60)) & ? [X61] : (p114(X61) & ~ p115(X61) & ~ p15(X61) & r1(X1, X61)))) & (~ p114(X1) | p115(X1) | (? [X62] : (p115(X62) & p16(X62) & r1(X1, X62)) & ? [X63] : (p115(X63) & ~ p16(X63) & r1(X1, X63))))) | ~ r1(X0, X1)) & ! [X64] : (p7(X64) | ~ r1(X0, X64))), inference(flattening, [], [f10])).
fof(f10, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25))))) & (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27))))) & (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29))))) & (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31))))) & (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33))))) & ((~ p100(X1) | p101(X1)) | (? [X34] : ((p101(X34) & ~ p102(X34) & p2(X34)) & r1(X1, X34)) & ? [X35] : ((p101(X35) & ~ p102(X35) & ~ p2(X35)) & r1(X1, X35)))) & ((~ p101(X1) | p102(X1)) | (? [X36] : ((p102(X36) & ~ p103(X36) & p3(X36)) & r1(X1, X36)) & ? [X37] : ((p102(X37) & ~ p103(X37) & ~ p3(X37)) & r1(X1, X37)))) & ((~ p102(X1) | p103(X1)) | (? [X38] : ((p103(X38) & ~ p104(X38) & p4(X38)) & r1(X1, X38)) & ? [X39] : ((p103(X39) & ~ p104(X39) & ~ p4(X39)) & r1(X1, X39)))) & ((~ p103(X1) | p104(X1)) | (? [X40] : ((p104(X40) & ~ p105(X40) & p5(X40)) & r1(X1, X40)) & ? [X41] : ((p104(X41) & ~ p105(X41) & ~ p5(X41)) & r1(X1, X41)))) & ((~ p104(X1) | p105(X1)) | (? [X42] : ((p105(X42) & ~ p106(X42) & p6(X42)) & r1(X1, X42)) & ? [X43] : ((p105(X43) & ~ p106(X43) & ~ p6(X43)) & r1(X1, X43)))) & ((~ p105(X1) | p106(X1)) | (? [X44] : ((p106(X44) & ~ p107(X44) & p7(X44)) & r1(X1, X44)) & ? [X45] : ((p106(X45) & ~ p107(X45) & ~ p7(X45)) & r1(X1, X45)))) & ((~ p106(X1) | p107(X1)) | (? [X46] : ((p107(X46) & ~ p108(X46) & p8(X46)) & r1(X1, X46)) & ? [X47] : ((p107(X47) & ~ p108(X47) & ~ p8(X47)) & r1(X1, X47)))) & ((~ p107(X1) | p108(X1)) | (? [X48] : ((p108(X48) & ~ p109(X48) & p9(X48)) & r1(X1, X48)) & ? [X49] : ((p108(X49) & ~ p109(X49) & ~ p9(X49)) & r1(X1, X49)))) & ((~ p108(X1) | p109(X1)) | (? [X50] : ((p109(X50) & ~ p110(X50) & p10(X50)) & r1(X1, X50)) & ? [X51] : ((p109(X51) & ~ p110(X51) & ~ p10(X51)) & r1(X1, X51)))) & ((~ p109(X1) | p110(X1)) | (? [X52] : ((p110(X52) & ~ p111(X52) & p11(X52)) & r1(X1, X52)) & ? [X53] : ((p110(X53) & ~ p111(X53) & ~ p11(X53)) & r1(X1, X53)))) & ((~ p110(X1) | p111(X1)) | (? [X54] : ((p111(X54) & ~ p112(X54) & p12(X54)) & r1(X1, X54)) & ? [X55] : ((p111(X55) & ~ p112(X55) & ~ p12(X55)) & r1(X1, X55)))) & ((~ p111(X1) | p112(X1)) | (? [X56] : ((p112(X56) & ~ p113(X56) & p13(X56)) & r1(X1, X56)) & ? [X57] : ((p112(X57) & ~ p113(X57) & ~ p13(X57)) & r1(X1, X57)))) & ((~ p112(X1) | p113(X1)) | (? [X58] : ((p113(X58) & ~ p114(X58) & p14(X58)) & r1(X1, X58)) & ? [X59] : ((p113(X59) & ~ p114(X59) & ~ p14(X59)) & r1(X1, X59)))) & ((~ p113(X1) | p114(X1)) | (? [X60] : ((p114(X60) & ~ p115(X60) & p15(X60)) & r1(X1, X60)) & ? [X61] : ((p114(X61) & ~ p115(X61) & ~ p15(X61)) & r1(X1, X61)))) & ((~ p114(X1) | p115(X1)) | (? [X62] : ((p115(X62) & p16(X62)) & r1(X1, X62)) & ? [X63] : ((p115(X63) & ~ p16(X63)) & r1(X1, X63))))) | ~ r1(X0, X1))) & ! [X64] : (p7(X64) | ~ r1(X0, X64))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25))))) & (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27))))) & (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29))))) & (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31))))) & (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X34] : (~ (p101(X34) & ~ p102(X34) & p2(X34)) | ~ r1(X1, X34)) & ~ ! [X35] : (~ (p101(X35) & ~ p102(X35) & ~ p2(X35)) | ~ r1(X1, X35)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X36] : (~ (p102(X36) & ~ p103(X36) & p3(X36)) | ~ r1(X1, X36)) & ~ ! [X37] : (~ (p102(X37) & ~ p103(X37) & ~ p3(X37)) | ~ r1(X1, X37)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X38] : (~ (p103(X38) & ~ p104(X38) & p4(X38)) | ~ r1(X1, X38)) & ~ ! [X39] : (~ (p103(X39) & ~ p104(X39) & ~ p4(X39)) | ~ r1(X1, X39)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X40] : (~ (p104(X40) & ~ p105(X40) & p5(X40)) | ~ r1(X1, X40)) & ~ ! [X41] : (~ (p104(X41) & ~ p105(X41) & ~ p5(X41)) | ~ r1(X1, X41)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X42] : (~ (p105(X42) & ~ p106(X42) & p6(X42)) | ~ r1(X1, X42)) & ~ ! [X43] : (~ (p105(X43) & ~ p106(X43) & ~ p6(X43)) | ~ r1(X1, X43)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X44] : (~ (p106(X44) & ~ p107(X44) & p7(X44)) | ~ r1(X1, X44)) & ~ ! [X45] : (~ (p106(X45) & ~ p107(X45) & ~ p7(X45)) | ~ r1(X1, X45)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X46] : (~ (p107(X46) & ~ p108(X46) & p8(X46)) | ~ r1(X1, X46)) & ~ ! [X47] : (~ (p107(X47) & ~ p108(X47) & ~ p8(X47)) | ~ r1(X1, X47)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X48] : (~ (p108(X48) & ~ p109(X48) & p9(X48)) | ~ r1(X1, X48)) & ~ ! [X49] : (~ (p108(X49) & ~ p109(X49) & ~ p9(X49)) | ~ r1(X1, X49)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X50] : (~ (p109(X50) & ~ p110(X50) & p10(X50)) | ~ r1(X1, X50)) & ~ ! [X51] : (~ (p109(X51) & ~ p110(X51) & ~ p10(X51)) | ~ r1(X1, X51)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X52] : (~ (p110(X52) & ~ p111(X52) & p11(X52)) | ~ r1(X1, X52)) & ~ ! [X53] : (~ (p110(X53) & ~ p111(X53) & ~ p11(X53)) | ~ r1(X1, X53)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X54] : (~ (p111(X54) & ~ p112(X54) & p12(X54)) | ~ r1(X1, X54)) & ~ ! [X55] : (~ (p111(X55) & ~ p112(X55) & ~ p12(X55)) | ~ r1(X1, X55)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X56] : (~ (p112(X56) & ~ p113(X56) & p13(X56)) | ~ r1(X1, X56)) & ~ ! [X57] : (~ (p112(X57) & ~ p113(X57) & ~ p13(X57)) | ~ r1(X1, X57)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X58] : (~ (p113(X58) & ~ p114(X58) & p14(X58)) | ~ r1(X1, X58)) & ~ ! [X59] : (~ (p113(X59) & ~ p114(X59) & ~ p14(X59)) | ~ r1(X1, X59)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X60] : (~ (p114(X60) & ~ p115(X60) & p15(X60)) | ~ r1(X1, X60)) & ~ ! [X61] : (~ (p114(X61) & ~ p115(X61) & ~ p15(X61)) | ~ r1(X1, X61)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X62] : (~ (p115(X62) & p16(X62)) | ~ r1(X1, X62)) & ~ ! [X63] : (~ (p115(X63) & ~ p16(X63)) | ~ r1(X1, X63))))) | ~ r1(X0, X1))) | ~ ! [X64] : (p7(X64) | ~ r1(X0, X64))), inference(pure_predicate_removal, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25))))) & (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27))))) & (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29))))) & (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31))))) & (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X34] : (~ (p101(X34) & ~ p102(X34) & p2(X34)) | ~ r1(X1, X34)) & ~ ! [X35] : (~ (p101(X35) & ~ p102(X35) & ~ p2(X35)) | ~ r1(X1, X35)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X36] : (~ (p102(X36) & ~ p103(X36) & p3(X36)) | ~ r1(X1, X36)) & ~ ! [X37] : (~ (p102(X37) & ~ p103(X37) & ~ p3(X37)) | ~ r1(X1, X37)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X38] : (~ (p103(X38) & ~ p104(X38) & p4(X38)) | ~ r1(X1, X38)) & ~ ! [X39] : (~ (p103(X39) & ~ p104(X39) & ~ p4(X39)) | ~ r1(X1, X39)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X40] : (~ (p104(X40) & ~ p105(X40) & p5(X40)) | ~ r1(X1, X40)) & ~ ! [X41] : (~ (p104(X41) & ~ p105(X41) & ~ p5(X41)) | ~ r1(X1, X41)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X42] : (~ (p105(X42) & ~ p106(X42) & p6(X42)) | ~ r1(X1, X42)) & ~ ! [X43] : (~ (p105(X43) & ~ p106(X43) & ~ p6(X43)) | ~ r1(X1, X43)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X44] : (~ (p106(X44) & ~ p107(X44) & p7(X44)) | ~ r1(X1, X44)) & ~ ! [X45] : (~ (p106(X45) & ~ p107(X45) & ~ p7(X45)) | ~ r1(X1, X45)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X46] : (~ (p107(X46) & ~ p108(X46) & p8(X46)) | ~ r1(X1, X46)) & ~ ! [X47] : (~ (p107(X47) & ~ p108(X47) & ~ p8(X47)) | ~ r1(X1, X47)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X48] : (~ (p108(X48) & ~ p109(X48) & p9(X48)) | ~ r1(X1, X48)) & ~ ! [X49] : (~ (p108(X49) & ~ p109(X49) & ~ p9(X49)) | ~ r1(X1, X49)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X50] : (~ (p109(X50) & ~ p110(X50) & p10(X50)) | ~ r1(X1, X50)) & ~ ! [X51] : (~ (p109(X51) & ~ p110(X51) & ~ p10(X51)) | ~ r1(X1, X51)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X52] : (~ (p110(X52) & ~ p111(X52) & p11(X52)) | ~ r1(X1, X52)) & ~ ! [X53] : (~ (p110(X53) & ~ p111(X53) & ~ p11(X53)) | ~ r1(X1, X53)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X54] : (~ (p111(X54) & ~ p112(X54) & p12(X54)) | ~ r1(X1, X54)) & ~ ! [X55] : (~ (p111(X55) & ~ p112(X55) & ~ p12(X55)) | ~ r1(X1, X55)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X56] : (~ (p112(X56) & ~ p113(X56) & p13(X56)) | ~ r1(X1, X56)) & ~ ! [X57] : (~ (p112(X57) & ~ p113(X57) & ~ p13(X57)) | ~ r1(X1, X57)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X58] : (~ (p113(X58) & ~ p114(X58) & p14(X58)) | ~ r1(X1, X58)) & ~ ! [X59] : (~ (p113(X59) & ~ p114(X59) & ~ p14(X59)) | ~ r1(X1, X59)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X60] : (~ (p114(X60) & ~ p115(X60) & p15(X60)) | ~ r1(X1, X60)) & ~ ! [X61] : (~ (p114(X61) & ~ p115(X61) & ~ p15(X61)) | ~ r1(X1, X61)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X62] : (~ (p115(X62) & ~ p116(X62) & p16(X62)) | ~ r1(X1, X62)) & ~ ! [X63] : (~ (p115(X63) & ~ p116(X63) & ~ p16(X63)) | ~ r1(X1, X63))))) | ~ r1(X0, X1))) | ~ ! [X64] : (p7(X64) | ~ r1(X0, X64))), inference(flattening, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25))))) & (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27))))) & (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29))))) & (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31))))) & (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X34] : (~ (p101(X34) & ~ p102(X34) & p2(X34)) | ~ r1(X1, X34)) & ~ ! [X35] : (~ (p101(X35) & ~ p102(X35) & ~ p2(X35)) | ~ r1(X1, X35)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X36] : (~ (p102(X36) & ~ p103(X36) & p3(X36)) | ~ r1(X1, X36)) & ~ ! [X37] : (~ (p102(X37) & ~ p103(X37) & ~ p3(X37)) | ~ r1(X1, X37)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X38] : (~ (p103(X38) & ~ p104(X38) & p4(X38)) | ~ r1(X1, X38)) & ~ ! [X39] : (~ (p103(X39) & ~ p104(X39) & ~ p4(X39)) | ~ r1(X1, X39)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X40] : (~ (p104(X40) & ~ p105(X40) & p5(X40)) | ~ r1(X1, X40)) & ~ ! [X41] : (~ (p104(X41) & ~ p105(X41) & ~ p5(X41)) | ~ r1(X1, X41)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X42] : (~ (p105(X42) & ~ p106(X42) & p6(X42)) | ~ r1(X1, X42)) & ~ ! [X43] : (~ (p105(X43) & ~ p106(X43) & ~ p6(X43)) | ~ r1(X1, X43)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X44] : (~ (p106(X44) & ~ p107(X44) & p7(X44)) | ~ r1(X1, X44)) & ~ ! [X45] : (~ (p106(X45) & ~ p107(X45) & ~ p7(X45)) | ~ r1(X1, X45)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X46] : (~ (p107(X46) & ~ p108(X46) & p8(X46)) | ~ r1(X1, X46)) & ~ ! [X47] : (~ (p107(X47) & ~ p108(X47) & ~ p8(X47)) | ~ r1(X1, X47)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X48] : (~ (p108(X48) & ~ p109(X48) & p9(X48)) | ~ r1(X1, X48)) & ~ ! [X49] : (~ (p108(X49) & ~ p109(X49) & ~ p9(X49)) | ~ r1(X1, X49)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X50] : (~ (p109(X50) & ~ p110(X50) & p10(X50)) | ~ r1(X1, X50)) & ~ ! [X51] : (~ (p109(X51) & ~ p110(X51) & ~ p10(X51)) | ~ r1(X1, X51)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X52] : (~ (p110(X52) & ~ p111(X52) & p11(X52)) | ~ r1(X1, X52)) & ~ ! [X53] : (~ (p110(X53) & ~ p111(X53) & ~ p11(X53)) | ~ r1(X1, X53)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X54] : (~ (p111(X54) & ~ p112(X54) & p12(X54)) | ~ r1(X1, X54)) & ~ ! [X55] : (~ (p111(X55) & ~ p112(X55) & ~ p12(X55)) | ~ r1(X1, X55)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X56] : (~ (p112(X56) & ~ p113(X56) & p13(X56)) | ~ r1(X1, X56)) & ~ ! [X57] : (~ (p112(X57) & ~ p113(X57) & ~ p13(X57)) | ~ r1(X1, X57)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X58] : (~ (p113(X58) & ~ p114(X58) & p14(X58)) | ~ r1(X1, X58)) & ~ ! [X59] : (~ (p113(X59) & ~ p114(X59) & ~ p14(X59)) | ~ r1(X1, X59)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X60] : (~ (p114(X60) & ~ p115(X60) & p15(X60)) | ~ r1(X1, X60)) & ~ ! [X61] : (~ (p114(X61) & ~ p115(X61) & ~ p15(X61)) | ~ r1(X1, X61)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X62] : (~ (p115(X62) & ~ p116(X62) & p16(X62)) | ~ r1(X1, X62)) & ~ ! [X63] : (~ (p115(X63) & ~ p116(X63) & ~ p16(X63)) | ~ r1(X1, X63))))) | ~ r1(X0, X1))) | ~ ! [X64] : (p7(X64) | ~ r1(X0, X64))), inference(rectify, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ p106(X1) | ((~ p7(X1) | ! [X0] : (~ p106(X0) | p7(X0) | ~ r1(X1, X0))) & (p7(X1) | ! [X0] : (~ p106(X0) | ~ p7(X0) | ~ r1(X1, X0))))) & (~ p107(X1) | ((~ p8(X1) | ! [X0] : (~ p107(X0) | p8(X0) | ~ r1(X1, X0))) & (p8(X1) | ! [X0] : (~ p107(X0) | ~ p8(X0) | ~ r1(X1, X0))))) & (~ p108(X1) | ((~ p9(X1) | ! [X0] : (~ p108(X0) | p9(X0) | ~ r1(X1, X0))) & (p9(X1) | ! [X0] : (~ p108(X0) | ~ p9(X0) | ~ r1(X1, X0))))) & (~ p109(X1) | ((~ p10(X1) | ! [X0] : (~ p109(X0) | p10(X0) | ~ r1(X1, X0))) & (p10(X1) | ! [X0] : (~ p109(X0) | ~ p10(X0) | ~ r1(X1, X0))))) & (~ p110(X1) | ((~ p11(X1) | ! [X0] : (~ p110(X0) | p11(X0) | ~ r1(X1, X0))) & (p11(X1) | ! [X0] : (~ p110(X0) | ~ p11(X0) | ~ r1(X1, X0))))) & (~ p111(X1) | ((~ p12(X1) | ! [X0] : (~ p111(X0) | p12(X0) | ~ r1(X1, X0))) & (p12(X1) | ! [X0] : (~ p111(X0) | ~ p12(X0) | ~ r1(X1, X0))))) & (~ p112(X1) | ((~ p13(X1) | ! [X0] : (~ p112(X0) | p13(X0) | ~ r1(X1, X0))) & (p13(X1) | ! [X0] : (~ p112(X0) | ~ p13(X0) | ~ r1(X1, X0))))) & (~ p113(X1) | ((~ p14(X1) | ! [X0] : (~ p113(X0) | p14(X0) | ~ r1(X1, X0))) & (p14(X1) | ! [X0] : (~ p113(X0) | ~ p14(X0) | ~ r1(X1, X0))))) & (~ p114(X1) | ((~ p15(X1) | ! [X0] : (~ p114(X0) | p15(X0) | ~ r1(X1, X0))) & (p15(X1) | ! [X0] : (~ p114(X0) | ~ p15(X0) | ~ r1(X1, X0))))) & (~ p115(X1) | ((~ p16(X1) | ! [X0] : (~ p115(X0) | p16(X0) | ~ r1(X1, X0))) & (p16(X1) | ! [X0] : (~ p115(X0) | ~ p16(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & p7(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & ~ p7(X0)) | ~ r1(X1, X0)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & p8(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & ~ p8(X0)) | ~ r1(X1, X0)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & p9(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & ~ p9(X0)) | ~ r1(X1, X0)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & p10(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & ~ p10(X0)) | ~ r1(X1, X0)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & p11(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & ~ p11(X0)) | ~ r1(X1, X0)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & p12(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & ~ p12(X0)) | ~ r1(X1, X0)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & p13(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & ~ p13(X0)) | ~ r1(X1, X0)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & p14(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & ~ p14(X0)) | ~ r1(X1, X0)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & p15(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & ~ p15(X0)) | ~ r1(X1, X0)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & p16(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & ~ p16(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p7(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ p106(X1) | ((~ p7(X1) | ! [X0] : (~ p106(X0) | p7(X0) | ~ r1(X1, X0))) & (p7(X1) | ! [X0] : (~ p106(X0) | ~ p7(X0) | ~ r1(X1, X0))))) & (~ p107(X1) | ((~ p8(X1) | ! [X0] : (~ p107(X0) | p8(X0) | ~ r1(X1, X0))) & (p8(X1) | ! [X0] : (~ p107(X0) | ~ p8(X0) | ~ r1(X1, X0))))) & (~ p108(X1) | ((~ p9(X1) | ! [X0] : (~ p108(X0) | p9(X0) | ~ r1(X1, X0))) & (p9(X1) | ! [X0] : (~ p108(X0) | ~ p9(X0) | ~ r1(X1, X0))))) & (~ p109(X1) | ((~ p10(X1) | ! [X0] : (~ p109(X0) | p10(X0) | ~ r1(X1, X0))) & (p10(X1) | ! [X0] : (~ p109(X0) | ~ p10(X0) | ~ r1(X1, X0))))) & (~ p110(X1) | ((~ p11(X1) | ! [X0] : (~ p110(X0) | p11(X0) | ~ r1(X1, X0))) & (p11(X1) | ! [X0] : (~ p110(X0) | ~ p11(X0) | ~ r1(X1, X0))))) & (~ p111(X1) | ((~ p12(X1) | ! [X0] : (~ p111(X0) | p12(X0) | ~ r1(X1, X0))) & (p12(X1) | ! [X0] : (~ p111(X0) | ~ p12(X0) | ~ r1(X1, X0))))) & (~ p112(X1) | ((~ p13(X1) | ! [X0] : (~ p112(X0) | p13(X0) | ~ r1(X1, X0))) & (p13(X1) | ! [X0] : (~ p112(X0) | ~ p13(X0) | ~ r1(X1, X0))))) & (~ p113(X1) | ((~ p14(X1) | ! [X0] : (~ p113(X0) | p14(X0) | ~ r1(X1, X0))) & (p14(X1) | ! [X0] : (~ p113(X0) | ~ p14(X0) | ~ r1(X1, X0))))) & (~ p114(X1) | ((~ p15(X1) | ! [X0] : (~ p114(X0) | p15(X0) | ~ r1(X1, X0))) & (p15(X1) | ! [X0] : (~ p114(X0) | ~ p15(X0) | ~ r1(X1, X0))))) & (~ p115(X1) | ((~ p16(X1) | ! [X0] : (~ p115(X0) | p16(X0) | ~ r1(X1, X0))) & (p16(X1) | ! [X0] : (~ p115(X0) | ~ p16(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & p7(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & ~ p7(X0)) | ~ r1(X1, X0)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & p8(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & ~ p8(X0)) | ~ r1(X1, X0)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & p9(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & ~ p9(X0)) | ~ r1(X1, X0)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & p10(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & ~ p10(X0)) | ~ r1(X1, X0)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & p11(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & ~ p11(X0)) | ~ r1(X1, X0)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & p12(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & ~ p12(X0)) | ~ r1(X1, X0)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & p13(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & ~ p13(X0)) | ~ r1(X1, X0)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & p14(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & ~ p14(X0)) | ~ r1(X1, X0)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & p15(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & ~ p15(X0)) | ~ r1(X1, X0)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & p16(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & ~ p16(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p7(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.015.p', main)).
fof(f13893, plain, (p101(sK62) | ~ r1(sK62, sK62) | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f13892, f375])).
fof(f375, plain, sP13(sK62), inference(resolution, [], [f360, f173])).
fof(f173, plain, ! [X0] : (~ sP31(X0) | sP13(X0)), inference(cnf_transformation, [], [f46])).
fof(f46, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p107(X0) | p106(X0)) & (~ p108(X0) | p107(X0)) & (~ p109(X0) | p108(X0)) & (~ p110(X0) | p109(X0)) & (~ p111(X0) | p110(X0)) & (~ p112(X0) | p111(X0)) & (~ p113(X0) | p112(X0)) & (~ p114(X0) | p113(X0)) & (~ p115(X0) | p114(X0)) & sP30(X0) & sP29(X0) & sP28(X0) & sP27(X0) & sP26(X0) & sP25(X0) & sP24(X0) & sP23(X0) & sP22(X0) & sP21(X0) & sP20(X0) & sP19(X0) & sP18(X0) & sP17(X0) & sP16(X0) & sP15(X0) & sP13(X0) & sP12(X0) & sP11(X0) & sP10(X0) & sP9(X0) & sP8(X0) & sP7(X0) & sP6(X0) & sP5(X0) & sP4(X0) & sP3(X0) & sP2(X0) & sP1(X0) & sP0(X0) & sP14(X0)) | ~ sP31(X0)), inference(rectify, [], [f45])).
fof(f45, plain, ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & sP30(X1) & sP29(X1) & sP28(X1) & sP27(X1) & sP26(X1) & sP25(X1) & sP24(X1) & sP23(X1) & sP22(X1) & sP21(X1) & sP20(X1) & sP19(X1) & sP18(X1) & sP17(X1) & sP16(X1) & sP15(X1) & sP13(X1) & sP12(X1) & sP11(X1) & sP10(X1) & sP9(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP4(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP14(X1)) | ~ sP31(X1)), inference(nnf_transformation, [], [f43])).
fof(f360, plain, sP31(sK62), inference(resolution, [], [f356, f157])).
fof(f356, plain, ! [X1] : (~ r1(sK62, X1) | sP31(X1)), inference(cnf_transformation, [], [f156])).
fof(f13892, plain, (~ sP13(sK62) | p101(sK62) | ~ r1(sK62, sK62) | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f13887, f358])).
fof(f358, plain, p100(sK62), inference(cnf_transformation, [], [f156])).
fof(f13887, plain, (~ p100(sK62) | ~ sP13(sK62) | p101(sK62) | ~ r1(sK62, sK62) | (spl63_44 | ~ spl63_45)), inference(resolution, [], [f13876, f621])).
fof(f621, plain, ! [X0, X1] : (r1(X1, sK34(X0)) | ~ p100(X0) | ~ sP13(X0) | p101(X0) | ~ r1(X1, X0)), inference(resolution, [], [f247, f158])).
fof(f158, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | r1(X0, X2) | ~ r1(X0, X1)), inference(cnf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1, X2] : (r1(X0, X2) | ~ r1(X1, X2) | ~ r1(X0, X1)), inference(flattening, [], [f8])).
fof(f8, plain, ! [X0, X1, X2] : (r1(X0, X2) | (~ r1(X1, X2) | ~ r1(X0, X1))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1, X2] : ((r1(X1, X2) & r1(X0, X1)) => r1(X0, X2)), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.015.p', transitivity)).
fof(f247, plain, ! [X0] : (r1(X0, sK34(X0)) | p101(X0) | ~ p100(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f88])).
fof(f88, plain, ! [X0] : (~ p100(X0) | p101(X0) | ((p101(sK34(X0)) & ~ p102(sK34(X0)) & p2(sK34(X0)) & r1(X0, sK34(X0))) & (p101(sK35(X0)) & ~ p102(sK35(X0)) & ~ p2(sK35(X0)) & r1(X0, sK35(X0)))) | ~ sP13(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK34, sK35])], [f85, f87, f86])).
fof(f86, plain, ! [X0] : (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) => (p101(sK34(X0)) & ~ p102(sK34(X0)) & p2(sK34(X0)) & r1(X0, sK34(X0)))), introduced(choice_axiom, [])).
fof(f87, plain, ! [X0] : (? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2)) => (p101(sK35(X0)) & ~ p102(sK35(X0)) & ~ p2(sK35(X0)) & r1(X0, sK35(X0)))), introduced(choice_axiom, [])).
fof(f85, plain, ! [X0] : (~ p100(X0) | p101(X0) | (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) & ? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2))) | ~ sP13(X0)), inference(rectify, [], [f84])).
fof(f84, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X34] : (p101(X34) & ~ p102(X34) & p2(X34) & r1(X1, X34)) & ? [X35] : (p101(X35) & ~ p102(X35) & ~ p2(X35) & r1(X1, X35))) | ~ sP13(X1)), inference(nnf_transformation, [], [f25])).
fof(f13876, plain, (~ r1(sK62, sK34(sK62)) | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f13875, f743])).
fof(f743, plain, (~ p102(sK34(sK62)) | spl63_44), inference(avatar_component_clause, [], [f742])).
fof(f742, plain, (spl63_44 <=> p102(sK34(sK62))), introduced(avatar_definition, [new_symbols(naming, [spl63_44])])).
fof(f13875, plain, (p102(sK34(sK62)) | ~ r1(sK62, sK34(sK62)) | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f13874, f644])).
fof(f644, plain, sP12(sK34(sK62)), inference(resolution, [], [f624, f172])).
fof(f172, plain, ! [X0] : (~ sP31(X0) | sP12(X0)), inference(cnf_transformation, [], [f46])).
fof(f624, plain, sP31(sK34(sK62)), inference(subsumption_resolution, [], [f623, f375])).
fof(f623, plain, (~ sP13(sK62) | sP31(sK34(sK62))), inference(subsumption_resolution, [], [f622, f358])).
fof(f622, plain, (~ p100(sK62) | ~ sP13(sK62) | sP31(sK34(sK62))), inference(subsumption_resolution, [], [f619, f357])).
fof(f619, plain, (p101(sK62) | ~ p100(sK62) | ~ sP13(sK62) | sP31(sK34(sK62))), inference(resolution, [], [f247, f356])).
fof(f13874, plain, (~ sP12(sK34(sK62)) | p102(sK34(sK62)) | ~ r1(sK62, sK34(sK62)) | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f13868, f749])).
fof(f749, plain, (p101(sK34(sK62)) | ~ spl63_45), inference(avatar_component_clause, [], [f747])).
fof(f747, plain, (spl63_45 <=> p101(sK34(sK62))), introduced(avatar_definition, [new_symbols(naming, [spl63_45])])).
fof(f13868, plain, (~ p101(sK34(sK62)) | ~ sP12(sK34(sK62)) | p102(sK34(sK62)) | ~ r1(sK62, sK34(sK62)) | (spl63_44 | ~ spl63_45)), inference(resolution, [], [f13860, f4324])).
fof(f4324, plain, (sP11(sK36(sK34(sK62))) | (spl63_44 | ~ spl63_45)), inference(resolution, [], [f4301, f171])).
fof(f171, plain, ! [X0] : (~ sP31(X0) | sP11(X0)), inference(cnf_transformation, [], [f46])).
fof(f4301, plain, (sP31(sK36(sK34(sK62))) | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f4300, f375])).
fof(f4300, plain, (sP31(sK36(sK34(sK62))) | ~ sP13(sK62) | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f4299, f358])).
fof(f4299, plain, (sP31(sK36(sK34(sK62))) | ~ p100(sK62) | ~ sP13(sK62) | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f4298, f357])).
fof(f4298, plain, (sP31(sK36(sK34(sK62))) | p101(sK62) | ~ p100(sK62) | ~ sP13(sK62) | (spl63_44 | ~ spl63_45)), inference(subsumption_resolution, [], [f4297, f749])).
fof(f4297, plain, (~ p101(sK34(sK62)) | sP31(sK36(sK34(sK62))) | p101(sK62) | ~ p100(sK62) | ~ sP13(sK62) | spl63_44), inference(subsumption_resolution, [], [f4296, f743])).
fof(f4296, plain, (p102(sK34(sK62)) | ~ p101(sK34(sK62)) | sP31(sK36(sK34(sK62))) | p101(sK62) | ~ p100(sK62) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f4240, f644])).
fof(f4240, plain, (~ sP12(sK34(sK62)) | p102(sK34(sK62)) | ~ p101(sK34(sK62)) | sP31(sK36(sK34(sK62))) | p101(sK62) | ~ p100(sK62) | ~ sP13(sK62)), inference(resolution, [], [f2107, f247])).
fof(f2107, plain, ! [X0] : (~ r1(sK62, X0) | ~ sP12(X0) | p102(X0) | ~ p101(X0) | sP31(sK36(X0))), inference(resolution, [], [f758, f356])).
fof(f758, plain, ! [X0, X1] : (r1(X1, sK36(X0)) | ~ p101(X0) | ~ sP12(X0) | p102(X0) | ~ r1(X1, X0)), inference(resolution, [], [f255, f158])).
fof(f255, plain, ! [X0] : (r1(X0, sK36(X0)) | p102(X0) | ~ p101(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ! [X0] : (~ p101(X0) | p102(X0) | ((p102(sK36(X0)) & ~ p103(sK36(X0)) & p3(sK36(X0)) & r1(X0, sK36(X0))) & (p102(sK37(X0)) & ~ p103(sK37(X0)) & ~ p3(sK37(X0)) & r1(X0, sK37(X0)))) | ~ sP12(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK36, sK37])], [f90, f92, f91])).
fof(f91, plain, ! [X0] : (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) => (p102(sK36(X0)) & ~ p103(sK36(X0)) & p3(sK36(X0)) & r1(X0, sK36(X0)))), introduced(choice_axiom, [])).
fof(f92, plain, ! [X0] : (? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2)) => (p102(sK37(X0)) & ~ p103(sK37(X0)) & ~ p3(sK37(X0)) & r1(X0, sK37(X0)))), introduced(choice_axiom, [])).
fof(f90, plain, ! [X0] : (~ p101(X0) | p102(X0) | (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) & ? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2))) | ~ sP12(X0)), inference(rectify, [], [f89])).
fof(f89, plain, ! [X1] : (~ p101(X1) | p102(X1) | (? [X36] : (p102(X36) & ~ p103(X36) & p3(X36) & r1(X1, X36)) & ? [X37] : (p102(X37) & ~ p103(X37) & ~ p3(X37) & r1(X1, X37))) | ~ sP12(X1)), inference(nnf_transformation, [], [f24])).
fof(f13860, plain, ! [X0] : (~ sP11(sK36(X0)) | ~ p101(X0) | ~ sP12(X0) | p102(X0) | ~ r1(sK62, X0)), inference(subsumption_resolution, [], [f13859, f758])).
fof(f13859, plain, ! [X0] : (~ sP11(sK36(X0)) | ~ r1(sK62, sK36(X0)) | ~ p101(X0) | ~ sP12(X0) | p102(X0) | ~ r1(sK62, X0)), inference(subsumption_resolution, [], [f13858, f257])).
fof(f257, plain, ! [X0] : (~ p103(sK36(X0)) | p102(X0) | ~ p101(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f93])).
fof(f13858, plain, ! [X0] : (~ sP11(sK36(X0)) | p103(sK36(X0)) | ~ r1(sK62, sK36(X0)) | ~ p101(X0) | ~ sP12(X0) | p102(X0) | ~ r1(sK62, X0)), inference(subsumption_resolution, [], [f13857, f258])).
fof(f258, plain, ! [X0] : (p102(sK36(X0)) | p102(X0) | ~ p101(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f93])).
fof(f13857, plain, ! [X0] : (~ p102(sK36(X0)) | ~ sP11(sK36(X0)) | p103(sK36(X0)) | ~ r1(sK62, sK36(X0)) | ~ p101(X0) | ~ sP12(X0) | p102(X0) | ~ r1(sK62, X0)), inference(duplicate_literal_removal, [], [f13854])).
fof(f13854, plain, ! [X0] : (~ p102(sK36(X0)) | ~ sP11(sK36(X0)) | p103(sK36(X0)) | ~ r1(sK62, sK36(X0)) | ~ p101(X0) | ~ sP12(X0) | p102(X0) | ~ r1(sK62, X0) | ~ sP11(sK36(X0))), inference(resolution, [], [f13850, f11869])).
fof(f11869, plain, ! [X11] : (sP10(sK38(sK36(X11))) | ~ p101(X11) | ~ sP12(X11) | p102(X11) | ~ r1(sK62, X11) | ~ sP11(sK36(X11))), inference(resolution, [], [f4851, f170])).
fof(f170, plain, ! [X0] : (~ sP31(X0) | sP10(X0)), inference(cnf_transformation, [], [f46])).
fof(f4851, plain, ! [X4] : (sP31(sK38(sK36(X4))) | ~ sP11(sK36(X4)) | ~ p101(X4) | ~ sP12(X4) | p102(X4) | ~ r1(sK62, X4)), inference(subsumption_resolution, [], [f4850, f258])).
fof(f4850, plain, ! [X4] : (~ sP11(sK36(X4)) | ~ p102(sK36(X4)) | sP31(sK38(sK36(X4))) | ~ p101(X4) | ~ sP12(X4) | p102(X4) | ~ r1(sK62, X4)), inference(subsumption_resolution, [], [f4799, f257])).
fof(f4799, plain, ! [X4] : (~ sP11(sK36(X4)) | p103(sK36(X4)) | ~ p102(sK36(X4)) | sP31(sK38(sK36(X4))) | ~ p101(X4) | ~ sP12(X4) | p102(X4) | ~ r1(sK62, X4)), inference(resolution, [], [f2182, f758])).
fof(f2182, plain, ! [X0] : (~ r1(sK62, X0) | ~ sP11(X0) | p103(X0) | ~ p102(X0) | sP31(sK38(X0))), inference(resolution, [], [f764, f356])).
fof(f764, plain, ! [X0, X1] : (r1(X1, sK38(X0)) | ~ p102(X0) | ~ sP11(X0) | p103(X0) | ~ r1(X1, X0)), inference(resolution, [], [f263, f158])).
fof(f263, plain, ! [X0] : (r1(X0, sK38(X0)) | p103(X0) | ~ p102(X0) | ~ sP11(X0)), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ! [X0] : (~ p102(X0) | p103(X0) | ((p103(sK38(X0)) & ~ p104(sK38(X0)) & p4(sK38(X0)) & r1(X0, sK38(X0))) & (p103(sK39(X0)) & ~ p104(sK39(X0)) & ~ p4(sK39(X0)) & r1(X0, sK39(X0)))) | ~ sP11(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK38, sK39])], [f95, f97, f96])).
fof(f96, plain, ! [X0] : (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) => (p103(sK38(X0)) & ~ p104(sK38(X0)) & p4(sK38(X0)) & r1(X0, sK38(X0)))), introduced(choice_axiom, [])).
fof(f97, plain, ! [X0] : (? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2)) => (p103(sK39(X0)) & ~ p104(sK39(X0)) & ~ p4(sK39(X0)) & r1(X0, sK39(X0)))), introduced(choice_axiom, [])).
fof(f95, plain, ! [X0] : (~ p102(X0) | p103(X0) | (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) & ? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2))) | ~ sP11(X0)), inference(rectify, [], [f94])).
fof(f94, plain, ! [X1] : (~ p102(X1) | p103(X1) | (? [X38] : (p103(X38) & ~ p104(X38) & p4(X38) & r1(X1, X38)) & ? [X39] : (p103(X39) & ~ p104(X39) & ~ p4(X39) & r1(X1, X39))) | ~ sP11(X1)), inference(nnf_transformation, [], [f23])).
fof(f13850, plain, ! [X0] : (~ sP10(sK38(X0)) | ~ p102(X0) | ~ sP11(X0) | p103(X0) | ~ r1(sK62, X0)), inference(subsumption_resolution, [], [f13849, f266])).
fof(f266, plain, ! [X0] : (p103(sK38(X0)) | p103(X0) | ~ p102(X0) | ~ sP11(X0)), inference(cnf_transformation, [], [f98])).
fof(f13849, plain, ! [X0] : (~ sP10(sK38(X0)) | ~ p103(sK38(X0)) | ~ p102(X0) | ~ sP11(X0) | p103(X0) | ~ r1(sK62, X0)), inference(subsumption_resolution, [], [f13848, f764])).
fof(f13848, plain, ! [X0] : (~ sP10(sK38(X0)) | ~ r1(sK62, sK38(X0)) | ~ p103(sK38(X0)) | ~ p102(X0) | ~ sP11(X0) | p103(X0) | ~ r1(sK62, X0)), inference(subsumption_resolution, [], [f13847, f265])).
fof(f265, plain, ! [X0] : (~ p104(sK38(X0)) | p103(X0) | ~ p102(X0) | ~ sP11(X0)), inference(cnf_transformation, [], [f98])).
fof(f13847, plain, ! [X0] : (~ sP10(sK38(X0)) | p104(sK38(X0)) | ~ r1(sK62, sK38(X0)) | ~ p103(sK38(X0)) | ~ p102(X0) | ~ sP11(X0) | p103(X0) | ~ r1(sK62, X0)), inference(duplicate_literal_removal, [], [f13844])).
fof(f13844, plain, ! [X0] : (~ sP10(sK38(X0)) | p104(sK38(X0)) | ~ r1(sK62, sK38(X0)) | ~ p103(sK38(X0)) | ~ p102(X0) | ~ sP11(X0) | p103(X0) | ~ r1(sK62, X0) | ~ sP10(sK38(X0))), inference(resolution, [], [f13843, f12052])).
fof(f12052, plain, ! [X10] : (sP9(sK40(sK38(X10))) | ~ p102(X10) | ~ sP11(X10) | p103(X10) | ~ r1(sK62, X10) | ~ sP10(sK38(X10))), inference(resolution, [], [f5111, f169])).
fof(f169, plain, ! [X0] : (~ sP31(X0) | sP9(X0)), inference(cnf_transformation, [], [f46])).
fof(f5111, plain, ! [X6] : (sP31(sK40(sK38(X6))) | ~ sP10(sK38(X6)) | ~ p102(X6) | ~ sP11(X6) | p103(X6) | ~ r1(sK62, X6)), inference(subsumption_resolution, [], [f5110, f266])).
fof(f5110, plain, ! [X6] : (~ sP10(sK38(X6)) | ~ p103(sK38(X6)) | sP31(sK40(sK38(X6))) | ~ p102(X6) | ~ sP11(X6) | p103(X6) | ~ r1(sK62, X6)), inference(subsumption_resolution, [], [f5063, f265])).
fof(f5063, plain, ! [X6] : (~ sP10(sK38(X6)) | p104(sK38(X6)) | ~ p103(sK38(X6)) | sP31(sK40(sK38(X6))) | ~ p102(X6) | ~ sP11(X6) | p103(X6) | ~ r1(sK62, X6)), inference(resolution, [], [f2257, f764])).
fof(f2257, plain, ! [X0] : (~ r1(sK62, X0) | ~ sP10(X0) | p104(X0) | ~ p103(X0) | sP31(sK40(X0))), inference(resolution, [], [f770, f356])).
fof(f770, plain, ! [X0, X1] : (r1(X1, sK40(X0)) | ~ p103(X0) | ~ sP10(X0) | p104(X0) | ~ r1(X1, X0)), inference(resolution, [], [f271, f158])).
fof(f271, plain, ! [X0] : (r1(X0, sK40(X0)) | p104(X0) | ~ p103(X0) | ~ sP10(X0)), inference(cnf_transformation, [], [f103])).
fof(f103, plain, ! [X0] : (~ p103(X0) | p104(X0) | ((p104(sK40(X0)) & ~ p105(sK40(X0)) & p5(sK40(X0)) & r1(X0, sK40(X0))) & (p104(sK41(X0)) & ~ p105(sK41(X0)) & ~ p5(sK41(X0)) & r1(X0, sK41(X0)))) | ~ sP10(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK40, sK41])], [f100, f102, f101])).
fof(f101, plain, ! [X0] : (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) => (p104(sK40(X0)) & ~ p105(sK40(X0)) & p5(sK40(X0)) & r1(X0, sK40(X0)))), introduced(choice_axiom, [])).
fof(f102, plain, ! [X0] : (? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2)) => (p104(sK41(X0)) & ~ p105(sK41(X0)) & ~ p5(sK41(X0)) & r1(X0, sK41(X0)))), introduced(choice_axiom, [])).
fof(f100, plain, ! [X0] : (~ p103(X0) | p104(X0) | (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) & ? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2))) | ~ sP10(X0)), inference(rectify, [], [f99])).
fof(f99, plain, ! [X1] : (~ p103(X1) | p104(X1) | (? [X40] : (p104(X40) & ~ p105(X40) & p5(X40) & r1(X1, X40)) & ? [X41] : (p104(X41) & ~ p105(X41) & ~ p5(X41) & r1(X1, X41))) | ~ sP10(X1)), inference(nnf_transformation, [], [f22])).
fof(f13843, plain, ! [X3] : (~ sP9(sK40(X3)) | ~ sP10(X3) | p104(X3) | ~ r1(sK62, X3) | ~ p103(X3)), inference(subsumption_resolution, [], [f13842, f770])).
fof(f13842, plain, ! [X3] : (~ p103(X3) | ~ sP10(X3) | p104(X3) | ~ r1(sK62, X3) | ~ sP9(sK40(X3)) | ~ r1(sK62, sK40(X3))), inference(subsumption_resolution, [], [f13841, f273])).
fof(f273, plain, ! [X0] : (~ p105(sK40(X0)) | p104(X0) | ~ p103(X0) | ~ sP10(X0)), inference(cnf_transformation, [], [f103])).
fof(f13841, plain, ! [X3] : (~ p103(X3) | ~ sP10(X3) | p104(X3) | ~ r1(sK62, X3) | ~ sP9(sK40(X3)) | p105(sK40(X3)) | ~ r1(sK62, sK40(X3))), inference(subsumption_resolution, [], [f13834, f274])).
fof(f274, plain, ! [X0] : (p104(sK40(X0)) | p104(X0) | ~ p103(X0) | ~ sP10(X0)), inference(cnf_transformation, [], [f103])).
fof(f13834, plain, ! [X3] : (~ p103(X3) | ~ sP10(X3) | p104(X3) | ~ r1(sK62, X3) | ~ sP9(sK40(X3)) | ~ p104(sK40(X3)) | p105(sK40(X3)) | ~ r1(sK62, sK40(X3))), inference(duplicate_literal_removal, [], [f13833])).
fof(f13833, plain, ! [X3] : (~ p103(X3) | ~ sP10(X3) | p104(X3) | ~ r1(sK62, X3) | ~ sP9(sK40(X3)) | ~ p104(sK40(X3)) | ~ sP9(sK40(X3)) | p105(sK40(X3)) | ~ r1(sK62, sK40(X3))), inference(resolution, [], [f12143, f3130])).
fof(f3130, plain, ! [X11] : (~ sP8(sK43(X11)) | ~ p104(X11) | ~ sP9(X11) | p105(X11) | ~ r1(sK62, X11)), inference(subsumption_resolution, [], [f3129, f278])).
fof(f278, plain, ! [X0] : (p105(sK43(X0)) | p105(X0) | ~ p104(X0) | ~ sP9(X0)), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ! [X0] : (~ p104(X0) | p105(X0) | ((p105(sK42(X0)) & ~ p106(sK42(X0)) & p6(sK42(X0)) & r1(X0, sK42(X0))) & (p105(sK43(X0)) & ~ p106(sK43(X0)) & ~ p6(sK43(X0)) & r1(X0, sK43(X0)))) | ~ sP9(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42, sK43])], [f105, f107, f106])).
fof(f106, plain, ! [X0] : (? [X1] : (p105(X1) & ~ p106(X1) & p6(X1) & r1(X0, X1)) => (p105(sK42(X0)) & ~ p106(sK42(X0)) & p6(sK42(X0)) & r1(X0, sK42(X0)))), introduced(choice_axiom, [])).
fof(f107, plain, ! [X0] : (? [X2] : (p105(X2) & ~ p106(X2) & ~ p6(X2) & r1(X0, X2)) => (p105(sK43(X0)) & ~ p106(sK43(X0)) & ~ p6(sK43(X0)) & r1(X0, sK43(X0)))), introduced(choice_axiom, [])).
fof(f105, plain, ! [X0] : (~ p104(X0) | p105(X0) | (? [X1] : (p105(X1) & ~ p106(X1) & p6(X1) & r1(X0, X1)) & ? [X2] : (p105(X2) & ~ p106(X2) & ~ p6(X2) & r1(X0, X2))) | ~ sP9(X0)), inference(rectify, [], [f104])).
fof(f104, plain, ! [X1] : (~ p104(X1) | p105(X1) | (? [X42] : (p105(X42) & ~ p106(X42) & p6(X42) & r1(X1, X42)) & ? [X43] : (p105(X43) & ~ p106(X43) & ~ p6(X43) & r1(X1, X43))) | ~ sP9(X1)), inference(nnf_transformation, [], [f21])).
fof(f3129, plain, ! [X11] : (~ sP8(sK43(X11)) | ~ p105(sK43(X11)) | ~ p104(X11) | ~ sP9(X11) | p105(X11) | ~ r1(sK62, X11)), inference(subsumption_resolution, [], [f3090, f277])).
fof(f277, plain, ! [X0] : (~ p106(sK43(X0)) | p105(X0) | ~ p104(X0) | ~ sP9(X0)), inference(cnf_transformation, [], [f108])).
fof(f3090, plain, ! [X11] : (~ sP8(sK43(X11)) | p106(sK43(X11)) | ~ p105(sK43(X11)) | ~ p104(X11) | ~ sP9(X11) | p105(X11) | ~ r1(sK62, X11)), inference(resolution, [], [f2405, f773])).
fof(f773, plain, ! [X0, X1] : (r1(X1, sK43(X0)) | ~ p104(X0) | ~ sP9(X0) | p105(X0) | ~ r1(X1, X0)), inference(resolution, [], [f275, f158])).
fof(f275, plain, ! [X0] : (r1(X0, sK43(X0)) | p105(X0) | ~ p104(X0) | ~ sP9(X0)), inference(cnf_transformation, [], [f108])).
fof(f2405, plain, ! [X1] : (~ r1(sK62, X1) | ~ sP8(X1) | p106(X1) | ~ p105(X1)), inference(subsumption_resolution, [], [f2371, f284])).
fof(f284, plain, ! [X0] : (~ p7(sK45(X0)) | p106(X0) | ~ p105(X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ! [X0] : (~ p105(X0) | p106(X0) | ((p106(sK44(X0)) & ~ p107(sK44(X0)) & p7(sK44(X0)) & r1(X0, sK44(X0))) & (p106(sK45(X0)) & ~ p107(sK45(X0)) & ~ p7(sK45(X0)) & r1(X0, sK45(X0)))) | ~ sP8(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK44, sK45])], [f110, f112, f111])).
fof(f111, plain, ! [X0] : (? [X1] : (p106(X1) & ~ p107(X1) & p7(X1) & r1(X0, X1)) => (p106(sK44(X0)) & ~ p107(sK44(X0)) & p7(sK44(X0)) & r1(X0, sK44(X0)))), introduced(choice_axiom, [])).
fof(f112, plain, ! [X0] : (? [X2] : (p106(X2) & ~ p107(X2) & ~ p7(X2) & r1(X0, X2)) => (p106(sK45(X0)) & ~ p107(sK45(X0)) & ~ p7(sK45(X0)) & r1(X0, sK45(X0)))), introduced(choice_axiom, [])).
fof(f110, plain, ! [X0] : (~ p105(X0) | p106(X0) | (? [X1] : (p106(X1) & ~ p107(X1) & p7(X1) & r1(X0, X1)) & ? [X2] : (p106(X2) & ~ p107(X2) & ~ p7(X2) & r1(X0, X2))) | ~ sP8(X0)), inference(rectify, [], [f109])).
fof(f109, plain, ! [X1] : (~ p105(X1) | p106(X1) | (? [X44] : (p106(X44) & ~ p107(X44) & p7(X44) & r1(X1, X44)) & ? [X45] : (p106(X45) & ~ p107(X45) & ~ p7(X45) & r1(X1, X45))) | ~ sP8(X1)), inference(nnf_transformation, [], [f20])).
fof(f2371, plain, ! [X1] : (~ p105(X1) | ~ sP8(X1) | p106(X1) | ~ r1(sK62, X1) | p7(sK45(X1))), inference(resolution, [], [f779, f355])).
fof(f355, plain, ! [X2] : (~ r1(sK62, X2) | p7(X2)), inference(cnf_transformation, [], [f156])).
fof(f779, plain, ! [X0, X1] : (r1(X1, sK45(X0)) | ~ p105(X0) | ~ sP8(X0) | p106(X0) | ~ r1(X1, X0)), inference(resolution, [], [f283, f158])).
fof(f283, plain, ! [X0] : (r1(X0, sK45(X0)) | p106(X0) | ~ p105(X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f113])).
fof(f12143, plain, ! [X9] : (sP8(sK43(sK40(X9))) | ~ p103(X9) | ~ sP10(X9) | p104(X9) | ~ r1(sK62, X9) | ~ sP9(sK40(X9))), inference(resolution, [], [f5241, f168])).
fof(f168, plain, ! [X0] : (~ sP31(X0) | sP8(X0)), inference(cnf_transformation, [], [f46])).
fof(f5241, plain, ! [X8] : (sP31(sK43(sK40(X8))) | ~ sP9(sK40(X8)) | ~ p103(X8) | ~ sP10(X8) | p104(X8) | ~ r1(sK62, X8)), inference(subsumption_resolution, [], [f5240, f274])).
fof(f5240, plain, ! [X8] : (~ sP9(sK40(X8)) | ~ p104(sK40(X8)) | sP31(sK43(sK40(X8))) | ~ p103(X8) | ~ sP10(X8) | p104(X8) | ~ r1(sK62, X8)), inference(subsumption_resolution, [], [f5197, f273])).
fof(f5197, plain, ! [X8] : (~ sP9(sK40(X8)) | p105(sK40(X8)) | ~ p104(sK40(X8)) | sP31(sK43(sK40(X8))) | ~ p103(X8) | ~ sP10(X8) | p104(X8) | ~ r1(sK62, X8)), inference(resolution, [], [f2295, f770])).
fof(f2295, plain, ! [X0] : (~ r1(sK62, X0) | ~ sP9(X0) | p105(X0) | ~ p104(X0) | sP31(sK43(X0))), inference(resolution, [], [f773, f356])).
fof(f3983, plain, ~ spl63_44, inference(avatar_contradiction_clause, [], [f3982])).
fof(f3982, plain, ($false | ~ spl63_44), inference(subsumption_resolution, [], [f3981, f375])).
fof(f3981, plain, (~ sP13(sK62) | ~ spl63_44), inference(subsumption_resolution, [], [f3980, f358])).
fof(f3980, plain, (~ p100(sK62) | ~ sP13(sK62) | ~ spl63_44), inference(subsumption_resolution, [], [f3979, f357])).
fof(f3979, plain, (p101(sK62) | ~ p100(sK62) | ~ sP13(sK62) | ~ spl63_44), inference(resolution, [], [f744, f249])).
fof(f249, plain, ! [X0] : (~ p102(sK34(X0)) | p101(X0) | ~ p100(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f88])).
fof(f744, plain, (p102(sK34(sK62)) | ~ spl63_44), inference(avatar_component_clause, [], [f742])).
fof(f3064, plain, spl63_45, inference(avatar_contradiction_clause, [], [f3063])).
fof(f3063, plain, ($false | spl63_45), inference(subsumption_resolution, [], [f3062, f375])).
fof(f3062, plain, (~ sP13(sK62) | spl63_45), inference(subsumption_resolution, [], [f3061, f358])).
fof(f3061, plain, (~ p100(sK62) | ~ sP13(sK62) | spl63_45), inference(subsumption_resolution, [], [f3060, f357])).
fof(f3060, plain, (p101(sK62) | ~ p100(sK62) | ~ sP13(sK62) | spl63_45), inference(resolution, [], [f748, f250])).
fof(f250, plain, ! [X0] : (p101(sK34(X0)) | p101(X0) | ~ p100(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f88])).
fof(f748, plain, (~ p101(sK34(sK62)) | spl63_45), inference(avatar_component_clause, [], [f747])).