fof(f6579, plain, $false, inference(avatar_sat_refutation, [], [f423, f441, f460, f480, f501, f522, f577, f579, f582, f586, f588, f591, f595, f600, f602, f605, f609, f614, f1117, f1278, f1428, f1824, f1999, f2019, f2552, f2562, f2563, f2648, f2654, f2823, f2907, f2921, f3094, f3106, f3108, f3110, f3118, f3130, f3150, f3168, f3180, f3261, f3345, f3356, f3382, f3476, f3526, f3548, f3563, f3565, f3577, f3583, f3672, f3674, f3690, f3693, f3694, f3777, f3846, f3848, f3850, f3907, f3919, f3922, f3924, f4015, f4021, f4057, f4142, f4184, f4230, f4232, f4244, f4249, f4251, f4252, f4254, f4361, f4363, f4371, f4460, f4473, f4475, f4477, f4478, f4494, f4495, f4577, f4673, f4824, f4832, f4834, f4838, f4844, f4946, f4970, f4994, f5019, f5043, f5089, f5145, f5151, f5203, f5221, f5246, f5247, f5367, f5377, f5384, f5456, f5470, f5489, f5496, f5574, f5594, f5694, f5698, f5700, f5702, f5778, f5780, f5932, f5946, f5973, f5980, f5981, f5982, f6124, f6128, f6196, f6238, f6260, f6261, f6423, f6578])).
fof(f6578, plain, (~ spl93_3 | ~ spl93_18), inference(avatar_contradiction_clause, [], [f6577])).
fof(f6577, plain, ($false | (~ spl93_3 | ~ spl93_18)), inference(subsumption_resolution, [], [f6576, f556])).
fof(f556, plain, sP33(sK92), inference(resolution, [], [f523, f266])).
fof(f266, plain, ! [X0] : (~ sP40(X0) | sP33(X0)), inference(cnf_transformation, [], [f50])).
fof(f50, plain, ! [X0] : (((~ p101(X0) | ~ p201(X0)) & (~ p101(X0) | ~ p301(X0)) & (~ p101(X0) | ~ p401(X0)) & (~ p101(X0) | ~ p501(X0)) & (~ p101(X0) | ~ p601(X0)) & (~ p201(X0) | ~ p301(X0)) & (~ p201(X0) | ~ p401(X0)) & (~ p201(X0) | ~ p501(X0)) & (~ p201(X0) | ~ p601(X0)) & (~ p301(X0) | ~ p401(X0)) & (~ p301(X0) | ~ p501(X0)) & (~ p301(X0) | ~ p601(X0)) & (~ p401(X0) | ~ p501(X0)) & (~ p401(X0) | ~ p601(X0)) & (~ p501(X0) | ~ p601(X0)) & sP39(X0) & sP38(X0) & sP37(X0) & sP36(X0) & sP35(X0) & (~ p202(X0) | ~ p302(X0)) & (~ p202(X0) | ~ p402(X0)) & (~ p202(X0) | ~ p502(X0)) & (~ p202(X0) | ~ p602(X0)) & (~ p302(X0) | ~ p402(X0)) & (~ p302(X0) | ~ p502(X0)) & (~ p302(X0) | ~ p602(X0)) & (~ p402(X0) | ~ p502(X0)) & (~ p402(X0) | ~ p602(X0)) & (~ p502(X0) | ~ p602(X0)) & sP9(X0) & sP34(X0) & sP33(X0) & sP32(X0) & sP31(X0) & sP30(X0) & sP29(X0) & sP28(X0) & sP27(X0) & (~ p303(X0) | ~ p403(X0)) & (~ p303(X0) | ~ p503(X0)) & (~ p303(X0) | ~ p603(X0)) & (~ p403(X0) | ~ p503(X0)) & (~ p403(X0) | ~ p603(X0)) & (~ p503(X0) | ~ p603(X0)) & sP8(X0) & sP7(X0) & sP26(X0) & sP25(X0) & sP24(X0) & sP6(X0) & sP23(X0) & sP22(X0) & sP21(X0) & sP20(X0) & sP19(X0) & sP18(X0) & (~ p404(X0) | ~ p504(X0)) & (~ p404(X0) | ~ p604(X0)) & (~ p504(X0) | ~ p604(X0)) & sP5(X0) & sP4(X0) & sP3(X0) & sP17(X0) & sP16(X0) & sP2(X0) & sP1(X0) & sP15(X0) & sP14(X0) & sP0(X0) & sP13(X0) & sP12(X0) & sP11(X0) & sP10(X0) & (~ p505(X0) | ~ p605(X0))) | ~ sP40(X0)), inference(rectify, [], [f49])).
fof(f49, plain, ! [X12] : (((~ p101(X12) | ~ p201(X12)) & (~ p101(X12) | ~ p301(X12)) & (~ p101(X12) | ~ p401(X12)) & (~ p101(X12) | ~ p501(X12)) & (~ p101(X12) | ~ p601(X12)) & (~ p201(X12) | ~ p301(X12)) & (~ p201(X12) | ~ p401(X12)) & (~ p201(X12) | ~ p501(X12)) & (~ p201(X12) | ~ p601(X12)) & (~ p301(X12) | ~ p401(X12)) & (~ p301(X12) | ~ p501(X12)) & (~ p301(X12) | ~ p601(X12)) & (~ p401(X12) | ~ p501(X12)) & (~ p401(X12) | ~ p601(X12)) & (~ p501(X12) | ~ p601(X12)) & sP39(X12) & sP38(X12) & sP37(X12) & sP36(X12) & sP35(X12) & (~ p202(X12) | ~ p302(X12)) & (~ p202(X12) | ~ p402(X12)) & (~ p202(X12) | ~ p502(X12)) & (~ p202(X12) | ~ p602(X12)) & (~ p302(X12) | ~ p402(X12)) & (~ p302(X12) | ~ p502(X12)) & (~ p302(X12) | ~ p602(X12)) & (~ p402(X12) | ~ p502(X12)) & (~ p402(X12) | ~ p602(X12)) & (~ p502(X12) | ~ p602(X12)) & sP9(X12) & sP34(X12) & sP33(X12) & sP32(X12) & sP31(X12) & sP30(X12) & sP29(X12) & sP28(X12) & sP27(X12) & (~ p303(X12) | ~ p403(X12)) & (~ p303(X12) | ~ p503(X12)) & (~ p303(X12) | ~ p603(X12)) & (~ p403(X12) | ~ p503(X12)) & (~ p403(X12) | ~ p603(X12)) & (~ p503(X12) | ~ p603(X12)) & sP8(X12) & sP7(X12) & sP26(X12) & sP25(X12) & sP24(X12) & sP6(X12) & sP23(X12) & sP22(X12) & sP21(X12) & sP20(X12) & sP19(X12) & sP18(X12) & (~ p404(X12) | ~ p504(X12)) & (~ p404(X12) | ~ p604(X12)) & (~ p504(X12) | ~ p604(X12)) & sP5(X12) & sP4(X12) & sP3(X12) & sP17(X12) & sP16(X12) & sP2(X12) & sP1(X12) & sP15(X12) & sP14(X12) & sP0(X12) & sP13(X12) & sP12(X12) & sP11(X12) & sP10(X12) & (~ p505(X12) | ~ p605(X12))) | ~ sP40(X12)), inference(nnf_transformation, [], [f47])).
fof(f47, plain, ! [X12] : (((~ p101(X12) | ~ p201(X12)) & (~ p101(X12) | ~ p301(X12)) & (~ p101(X12) | ~ p401(X12)) & (~ p101(X12) | ~ p501(X12)) & (~ p101(X12) | ~ p601(X12)) & (~ p201(X12) | ~ p301(X12)) & (~ p201(X12) | ~ p401(X12)) & (~ p201(X12) | ~ p501(X12)) & (~ p201(X12) | ~ p601(X12)) & (~ p301(X12) | ~ p401(X12)) & (~ p301(X12) | ~ p501(X12)) & (~ p301(X12) | ~ p601(X12)) & (~ p401(X12) | ~ p501(X12)) & (~ p401(X12) | ~ p601(X12)) & (~ p501(X12) | ~ p601(X12)) & sP39(X12) & sP38(X12) & sP37(X12) & sP36(X12) & sP35(X12) & (~ p202(X12) | ~ p302(X12)) & (~ p202(X12) | ~ p402(X12)) & (~ p202(X12) | ~ p502(X12)) & (~ p202(X12) | ~ p602(X12)) & (~ p302(X12) | ~ p402(X12)) & (~ p302(X12) | ~ p502(X12)) & (~ p302(X12) | ~ p602(X12)) & (~ p402(X12) | ~ p502(X12)) & (~ p402(X12) | ~ p602(X12)) & (~ p502(X12) | ~ p602(X12)) & sP9(X12) & sP34(X12) & sP33(X12) & sP32(X12) & sP31(X12) & sP30(X12) & sP29(X12) & sP28(X12) & sP27(X12) & (~ p303(X12) | ~ p403(X12)) & (~ p303(X12) | ~ p503(X12)) & (~ p303(X12) | ~ p603(X12)) & (~ p403(X12) | ~ p503(X12)) & (~ p403(X12) | ~ p603(X12)) & (~ p503(X12) | ~ p603(X12)) & sP8(X12) & sP7(X12) & sP26(X12) & sP25(X12) & sP24(X12) & sP6(X12) & sP23(X12) & sP22(X12) & sP21(X12) & sP20(X12) & sP19(X12) & sP18(X12) & (~ p404(X12) | ~ p504(X12)) & (~ p404(X12) | ~ p604(X12)) & (~ p504(X12) | ~ p604(X12)) & sP5(X12) & sP4(X12) & sP3(X12) & sP17(X12) & sP16(X12) & sP2(X12) & sP1(X12) & sP15(X12) & sP14(X12) & sP0(X12) & sP13(X12) & sP12(X12) & sP11(X12) & sP10(X12) & (~ p505(X12) | ~ p605(X12))) | ~ sP40(X12)), inference(usedef, [], [e47])).
fof(e47, plain, ! [X12] : (sP40(X12) <=> ((~ p101(X12) | ~ p201(X12)) & (~ p101(X12) | ~ p301(X12)) & (~ p101(X12) | ~ p401(X12)) & (~ p101(X12) | ~ p501(X12)) & (~ p101(X12) | ~ p601(X12)) & (~ p201(X12) | ~ p301(X12)) & (~ p201(X12) | ~ p401(X12)) & (~ p201(X12) | ~ p501(X12)) & (~ p201(X12) | ~ p601(X12)) & (~ p301(X12) | ~ p401(X12)) & (~ p301(X12) | ~ p501(X12)) & (~ p301(X12) | ~ p601(X12)) & (~ p401(X12) | ~ p501(X12)) & (~ p401(X12) | ~ p601(X12)) & (~ p501(X12) | ~ p601(X12)) & sP39(X12) & sP38(X12) & sP37(X12) & sP36(X12) & sP35(X12) & (~ p202(X12) | ~ p302(X12)) & (~ p202(X12) | ~ p402(X12)) & (~ p202(X12) | ~ p502(X12)) & (~ p202(X12) | ~ p602(X12)) & (~ p302(X12) | ~ p402(X12)) & (~ p302(X12) | ~ p502(X12)) & (~ p302(X12) | ~ p602(X12)) & (~ p402(X12) | ~ p502(X12)) & (~ p402(X12) | ~ p602(X12)) & (~ p502(X12) | ~ p602(X12)) & sP9(X12) & sP34(X12) & sP33(X12) & sP32(X12) & sP31(X12) & sP30(X12) & sP29(X12) & sP28(X12) & sP27(X12) & (~ p303(X12) | ~ p403(X12)) & (~ p303(X12) | ~ p503(X12)) & (~ p303(X12) | ~ p603(X12)) & (~ p403(X12) | ~ p503(X12)) & (~ p403(X12) | ~ p603(X12)) & (~ p503(X12) | ~ p603(X12)) & sP8(X12) & sP7(X12) & sP26(X12) & sP25(X12) & sP24(X12) & sP6(X12) & sP23(X12) & sP22(X12) & sP21(X12) & sP20(X12) & sP19(X12) & sP18(X12) & (~ p404(X12) | ~ p504(X12)) & (~ p404(X12) | ~ p604(X12)) & (~ p504(X12) | ~ p604(X12)) & sP5(X12) & sP4(X12) & sP3(X12) & sP17(X12) & sP16(X12) & sP2(X12) & sP1(X12) & sP15(X12) & sP14(X12) & sP0(X12) & sP13(X12) & sP12(X12) & sP11(X12) & sP10(X12) & (~ p505(X12) | ~ p605(X12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f523, plain, sP40(sK92), inference(resolution, [], [f399, f400])).
fof(f400, plain, r1(sK91, sK92), inference(cnf_transformation, [], [f223])).
fof(f223, plain, (((p101(sK92) | ! [X2] : (p102(X2) | ~ r1(sK92, X2)) | ! [X3] : (p103(X3) | ~ r1(sK92, X3)) | ! [X4] : (p104(X4) | ~ r1(sK92, X4)) | ! [X5] : (p105(X5) | ~ r1(sK92, X5))) & (p201(sK92) | p202(sK92) | ! [X6] : (p203(X6) | ~ r1(sK92, X6)) | ! [X7] : (p204(X7) | ~ r1(sK92, X7)) | ! [X8] : (p205(X8) | ~ r1(sK92, X8))) & (p301(sK92) | p302(sK92) | p303(sK92) | ! [X9] : (p304(X9) | ~ r1(sK92, X9)) | ! [X10] : (p305(X10) | ~ r1(sK92, X10))) & (p401(sK92) | p402(sK92) | p403(sK92) | p404(sK92) | ! [X11] : (p405(X11) | ~ r1(sK92, X11))) & (p501(sK92) | p502(sK92) | p503(sK92) | p504(sK92) | p505(sK92)) & (p601(sK92) | p602(sK92) | p603(sK92) | p604(sK92) | p605(sK92)) & r1(sK91, sK92)) & ! [X12] : (sP40(X12) | ~ r1(sK91, X12))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK91, sK92])], [f48, f222, f221])).
fof(f221, plain, (? [X0] : (? [X1] : ((p101(X1) | ! [X2] : (p102(X2) | ~ r1(X1, X2)) | ! [X3] : (p103(X3) | ~ r1(X1, X3)) | ! [X4] : (p104(X4) | ~ r1(X1, X4)) | ! [X5] : (p105(X5) | ~ r1(X1, X5))) & (p201(X1) | p202(X1) | ! [X6] : (p203(X6) | ~ r1(X1, X6)) | ! [X7] : (p204(X7) | ~ r1(X1, X7)) | ! [X8] : (p205(X8) | ~ r1(X1, X8))) & (p301(X1) | p302(X1) | p303(X1) | ! [X9] : (p304(X9) | ~ r1(X1, X9)) | ! [X10] : (p305(X10) | ~ r1(X1, X10))) & (p401(X1) | p402(X1) | p403(X1) | p404(X1) | ! [X11] : (p405(X11) | ~ r1(X1, X11))) & (p501(X1) | p502(X1) | p503(X1) | p504(X1) | p505(X1)) & (p601(X1) | p602(X1) | p603(X1) | p604(X1) | p605(X1)) & r1(X0, X1)) & ! [X12] : (sP40(X12) | ~ r1(X0, X12))) => (? [X1] : ((p101(X1) | ! [X2] : (p102(X2) | ~ r1(X1, X2)) | ! [X3] : (p103(X3) | ~ r1(X1, X3)) | ! [X4] : (p104(X4) | ~ r1(X1, X4)) | ! [X5] : (p105(X5) | ~ r1(X1, X5))) & (p201(X1) | p202(X1) | ! [X6] : (p203(X6) | ~ r1(X1, X6)) | ! [X7] : (p204(X7) | ~ r1(X1, X7)) | ! [X8] : (p205(X8) | ~ r1(X1, X8))) & (p301(X1) | p302(X1) | p303(X1) | ! [X9] : (p304(X9) | ~ r1(X1, X9)) | ! [X10] : (p305(X10) | ~ r1(X1, X10))) & (p401(X1) | p402(X1) | p403(X1) | p404(X1) | ! [X11] : (p405(X11) | ~ r1(X1, X11))) & (p501(X1) | p502(X1) | p503(X1) | p504(X1) | p505(X1)) & (p601(X1) | p602(X1) | p603(X1) | p604(X1) | p605(X1)) & r1(sK91, X1)) & ! [X12] : (sP40(X12) | ~ r1(sK91, X12)))), introduced(choice_axiom, [])).
fof(f222, plain, (? [X1] : ((p101(X1) | ! [X2] : (p102(X2) | ~ r1(X1, X2)) | ! [X3] : (p103(X3) | ~ r1(X1, X3)) | ! [X4] : (p104(X4) | ~ r1(X1, X4)) | ! [X5] : (p105(X5) | ~ r1(X1, X5))) & (p201(X1) | p202(X1) | ! [X6] : (p203(X6) | ~ r1(X1, X6)) | ! [X7] : (p204(X7) | ~ r1(X1, X7)) | ! [X8] : (p205(X8) | ~ r1(X1, X8))) & (p301(X1) | p302(X1) | p303(X1) | ! [X9] : (p304(X9) | ~ r1(X1, X9)) | ! [X10] : (p305(X10) | ~ r1(X1, X10))) & (p401(X1) | p402(X1) | p403(X1) | p404(X1) | ! [X11] : (p405(X11) | ~ r1(X1, X11))) & (p501(X1) | p502(X1) | p503(X1) | p504(X1) | p505(X1)) & (p601(X1) | p602(X1) | p603(X1) | p604(X1) | p605(X1)) & r1(sK91, X1)) => ((p101(sK92) | ! [X2] : (p102(X2) | ~ r1(sK92, X2)) | ! [X3] : (p103(X3) | ~ r1(sK92, X3)) | ! [X4] : (p104(X4) | ~ r1(sK92, X4)) | ! [X5] : (p105(X5) | ~ r1(sK92, X5))) & (p201(sK92) | p202(sK92) | ! [X6] : (p203(X6) | ~ r1(sK92, X6)) | ! [X7] : (p204(X7) | ~ r1(sK92, X7)) | ! [X8] : (p205(X8) | ~ r1(sK92, X8))) & (p301(sK92) | p302(sK92) | p303(sK92) | ! [X9] : (p304(X9) | ~ r1(sK92, X9)) | ! [X10] : (p305(X10) | ~ r1(sK92, X10))) & (p401(sK92) | p402(sK92) | p403(sK92) | p404(sK92) | ! [X11] : (p405(X11) | ~ r1(sK92, X11))) & (p501(sK92) | p502(sK92) | p503(sK92) | p504(sK92) | p505(sK92)) & (p601(sK92) | p602(sK92) | p603(sK92) | p604(sK92) | p605(sK92)) & r1(sK91, sK92))), introduced(choice_axiom, [])).
fof(f48, plain, ? [X0] : (? [X1] : ((p101(X1) | ! [X2] : (p102(X2) | ~ r1(X1, X2)) | ! [X3] : (p103(X3) | ~ r1(X1, X3)) | ! [X4] : (p104(X4) | ~ r1(X1, X4)) | ! [X5] : (p105(X5) | ~ r1(X1, X5))) & (p201(X1) | p202(X1) | ! [X6] : (p203(X6) | ~ r1(X1, X6)) | ! [X7] : (p204(X7) | ~ r1(X1, X7)) | ! [X8] : (p205(X8) | ~ r1(X1, X8))) & (p301(X1) | p302(X1) | p303(X1) | ! [X9] : (p304(X9) | ~ r1(X1, X9)) | ! [X10] : (p305(X10) | ~ r1(X1, X10))) & (p401(X1) | p402(X1) | p403(X1) | p404(X1) | ! [X11] : (p405(X11) | ~ r1(X1, X11))) & (p501(X1) | p502(X1) | p503(X1) | p504(X1) | p505(X1)) & (p601(X1) | p602(X1) | p603(X1) | p604(X1) | p605(X1)) & r1(X0, X1)) & ! [X12] : (sP40(X12) | ~ r1(X0, X12))), inference(definition_folding, [], [f6, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7])).
fof(f7, plain, ! [X12] : (? [X57] : (~ p305(X57) & r1(X12, X57)) | ? [X58] : (~ p405(X58) & r1(X12, X58)) | ~ sP0(X12)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X12] : (sP0(X12) <=> (? [X57] : (~ p305(X57) & r1(X12, X57)) | ? [X58] : (~ p405(X58) & r1(X12, X58)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f8, plain, ! [X12] : (? [X53] : (~ p205(X53) & r1(X12, X53)) | ? [X54] : (~ p405(X54) & r1(X12, X54)) | ~ sP1(X12)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X12] : (sP1(X12) <=> (? [X53] : (~ p205(X53) & r1(X12, X53)) | ? [X54] : (~ p405(X54) & r1(X12, X54)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f9, plain, ! [X12] : (? [X51] : (~ p205(X51) & r1(X12, X51)) | ? [X52] : (~ p305(X52) & r1(X12, X52)) | ~ sP2(X12)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X12] : (sP2(X12) <=> (? [X51] : (~ p205(X51) & r1(X12, X51)) | ? [X52] : (~ p305(X52) & r1(X12, X52)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f10, plain, ! [X12] : (? [X47] : (~ p105(X47) & r1(X12, X47)) | ? [X48] : (~ p405(X48) & r1(X12, X48)) | ~ sP3(X12)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X12] : (sP3(X12) <=> (? [X47] : (~ p105(X47) & r1(X12, X47)) | ? [X48] : (~ p405(X48) & r1(X12, X48)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f11, plain, ! [X12] : (? [X45] : (~ p105(X45) & r1(X12, X45)) | ? [X46] : (~ p305(X46) & r1(X12, X46)) | ~ sP4(X12)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X12] : (sP4(X12) <=> (? [X45] : (~ p105(X45) & r1(X12, X45)) | ? [X46] : (~ p305(X46) & r1(X12, X46)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f12, plain, ! [X12] : (? [X43] : (~ p105(X43) & r1(X12, X43)) | ? [X44] : (~ p205(X44) & r1(X12, X44)) | ~ sP5(X12)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X12] : (sP5(X12) <=> (? [X43] : (~ p105(X43) & r1(X12, X43)) | ? [X44] : (~ p205(X44) & r1(X12, X44)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f13, plain, ! [X12] : (? [X35] : (~ p204(X35) & r1(X12, X35)) | ? [X36] : (~ p304(X36) & r1(X12, X36)) | ~ sP6(X12)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X12] : (sP6(X12) <=> (? [X35] : (~ p204(X35) & r1(X12, X35)) | ? [X36] : (~ p304(X36) & r1(X12, X36)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f14, plain, ! [X12] : (? [X30] : (~ p104(X30) & r1(X12, X30)) | ? [X31] : (~ p304(X31) & r1(X12, X31)) | ~ sP7(X12)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X12] : (sP7(X12) <=> (? [X30] : (~ p104(X30) & r1(X12, X30)) | ? [X31] : (~ p304(X31) & r1(X12, X31)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f15, plain, ! [X12] : (? [X28] : (~ p104(X28) & r1(X12, X28)) | ? [X29] : (~ p204(X29) & r1(X12, X29)) | ~ sP8(X12)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X12] : (sP8(X12) <=> (? [X28] : (~ p104(X28) & r1(X12, X28)) | ? [X29] : (~ p204(X29) & r1(X12, X29)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f16, plain, ! [X12] : (? [X18] : (~ p103(X18) & r1(X12, X18)) | ? [X19] : (~ p203(X19) & r1(X12, X19)) | ~ sP9(X12)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X12] : (sP9(X12) <=> (? [X18] : (~ p103(X18) & r1(X12, X18)) | ? [X19] : (~ p203(X19) & r1(X12, X19)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f17, plain, ! [X12] : (? [X62] : (~ p405(X62) & r1(X12, X62)) | ~ p605(X12) | ~ sP10(X12)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X12] : (sP10(X12) <=> (? [X62] : (~ p405(X62) & r1(X12, X62)) | ~ p605(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f18, plain, ! [X12] : (? [X61] : (~ p405(X61) & r1(X12, X61)) | ~ p505(X12) | ~ sP11(X12)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X12] : (sP11(X12) <=> (? [X61] : (~ p405(X61) & r1(X12, X61)) | ~ p505(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f19, plain, ! [X12] : (? [X60] : (~ p305(X60) & r1(X12, X60)) | ~ p605(X12) | ~ sP12(X12)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X12] : (sP12(X12) <=> (? [X60] : (~ p305(X60) & r1(X12, X60)) | ~ p605(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f20, plain, ! [X12] : (? [X59] : (~ p305(X59) & r1(X12, X59)) | ~ p505(X12) | ~ sP13(X12)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X12] : (sP13(X12) <=> (? [X59] : (~ p305(X59) & r1(X12, X59)) | ~ p505(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f21, plain, ! [X12] : (? [X56] : (~ p205(X56) & r1(X12, X56)) | ~ p605(X12) | ~ sP14(X12)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X12] : (sP14(X12) <=> (? [X56] : (~ p205(X56) & r1(X12, X56)) | ~ p605(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f22, plain, ! [X12] : (? [X55] : (~ p205(X55) & r1(X12, X55)) | ~ p505(X12) | ~ sP15(X12)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X12] : (sP15(X12) <=> (? [X55] : (~ p205(X55) & r1(X12, X55)) | ~ p505(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f23, plain, ! [X12] : (? [X50] : (~ p105(X50) & r1(X12, X50)) | ~ p605(X12) | ~ sP16(X12)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X12] : (sP16(X12) <=> (? [X50] : (~ p105(X50) & r1(X12, X50)) | ~ p605(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f24, plain, ! [X12] : (? [X49] : (~ p105(X49) & r1(X12, X49)) | ~ p505(X12) | ~ sP17(X12)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X12] : (sP17(X12) <=> (? [X49] : (~ p105(X49) & r1(X12, X49)) | ~ p505(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f25, plain, ! [X12] : (? [X42] : (~ p304(X42) & r1(X12, X42)) | ~ p604(X12) | ~ sP18(X12)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X12] : (sP18(X12) <=> (? [X42] : (~ p304(X42) & r1(X12, X42)) | ~ p604(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f26, plain, ! [X12] : (? [X41] : (~ p304(X41) & r1(X12, X41)) | ~ p504(X12) | ~ sP19(X12)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X12] : (sP19(X12) <=> (? [X41] : (~ p304(X41) & r1(X12, X41)) | ~ p504(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f27, plain, ! [X12] : (? [X40] : (~ p304(X40) & r1(X12, X40)) | ~ p404(X12) | ~ sP20(X12)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X12] : (sP20(X12) <=> (? [X40] : (~ p304(X40) & r1(X12, X40)) | ~ p404(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f28, plain, ! [X12] : (? [X39] : (~ p204(X39) & r1(X12, X39)) | ~ p604(X12) | ~ sP21(X12)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X12] : (sP21(X12) <=> (? [X39] : (~ p204(X39) & r1(X12, X39)) | ~ p604(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f29, plain, ! [X12] : (? [X38] : (~ p204(X38) & r1(X12, X38)) | ~ p504(X12) | ~ sP22(X12)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X12] : (sP22(X12) <=> (? [X38] : (~ p204(X38) & r1(X12, X38)) | ~ p504(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f30, plain, ! [X12] : (? [X37] : (~ p204(X37) & r1(X12, X37)) | ~ p404(X12) | ~ sP23(X12)), inference(usedef, [], [e30])).
fof(e30, plain, ! [X12] : (sP23(X12) <=> (? [X37] : (~ p204(X37) & r1(X12, X37)) | ~ p404(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f31, plain, ! [X12] : (? [X34] : (~ p104(X34) & r1(X12, X34)) | ~ p604(X12) | ~ sP24(X12)), inference(usedef, [], [e31])).
fof(e31, plain, ! [X12] : (sP24(X12) <=> (? [X34] : (~ p104(X34) & r1(X12, X34)) | ~ p604(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f32, plain, ! [X12] : (? [X33] : (~ p104(X33) & r1(X12, X33)) | ~ p504(X12) | ~ sP25(X12)), inference(usedef, [], [e32])).
fof(e32, plain, ! [X12] : (sP25(X12) <=> (? [X33] : (~ p104(X33) & r1(X12, X33)) | ~ p504(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f33, plain, ! [X12] : (? [X32] : (~ p104(X32) & r1(X12, X32)) | ~ p404(X12) | ~ sP26(X12)), inference(usedef, [], [e33])).
fof(e33, plain, ! [X12] : (sP26(X12) <=> (? [X32] : (~ p104(X32) & r1(X12, X32)) | ~ p404(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f34, plain, ! [X12] : (? [X27] : (~ p203(X27) & r1(X12, X27)) | ~ p603(X12) | ~ sP27(X12)), inference(usedef, [], [e34])).
fof(e34, plain, ! [X12] : (sP27(X12) <=> (? [X27] : (~ p203(X27) & r1(X12, X27)) | ~ p603(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f35, plain, ! [X12] : (? [X26] : (~ p203(X26) & r1(X12, X26)) | ~ p503(X12) | ~ sP28(X12)), inference(usedef, [], [e35])).
fof(e35, plain, ! [X12] : (sP28(X12) <=> (? [X26] : (~ p203(X26) & r1(X12, X26)) | ~ p503(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f36, plain, ! [X12] : (? [X25] : (~ p203(X25) & r1(X12, X25)) | ~ p403(X12) | ~ sP29(X12)), inference(usedef, [], [e36])).
fof(e36, plain, ! [X12] : (sP29(X12) <=> (? [X25] : (~ p203(X25) & r1(X12, X25)) | ~ p403(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f37, plain, ! [X12] : (? [X24] : (~ p203(X24) & r1(X12, X24)) | ~ p303(X12) | ~ sP30(X12)), inference(usedef, [], [e37])).
fof(e37, plain, ! [X12] : (sP30(X12) <=> (? [X24] : (~ p203(X24) & r1(X12, X24)) | ~ p303(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f38, plain, ! [X12] : (? [X23] : (~ p103(X23) & r1(X12, X23)) | ~ p603(X12) | ~ sP31(X12)), inference(usedef, [], [e38])).
fof(e38, plain, ! [X12] : (sP31(X12) <=> (? [X23] : (~ p103(X23) & r1(X12, X23)) | ~ p603(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f39, plain, ! [X12] : (? [X22] : (~ p103(X22) & r1(X12, X22)) | ~ p503(X12) | ~ sP32(X12)), inference(usedef, [], [e39])).
fof(e39, plain, ! [X12] : (sP32(X12) <=> (? [X22] : (~ p103(X22) & r1(X12, X22)) | ~ p503(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f40, plain, ! [X12] : (? [X21] : (~ p103(X21) & r1(X12, X21)) | ~ p403(X12) | ~ sP33(X12)), inference(usedef, [], [e40])).
fof(e40, plain, ! [X12] : (sP33(X12) <=> (? [X21] : (~ p103(X21) & r1(X12, X21)) | ~ p403(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f41, plain, ! [X12] : (? [X20] : (~ p103(X20) & r1(X12, X20)) | ~ p303(X12) | ~ sP34(X12)), inference(usedef, [], [e41])).
fof(e41, plain, ! [X12] : (sP34(X12) <=> (? [X20] : (~ p103(X20) & r1(X12, X20)) | ~ p303(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f42, plain, ! [X12] : (? [X17] : (~ p102(X17) & r1(X12, X17)) | ~ p602(X12) | ~ sP35(X12)), inference(usedef, [], [e42])).
fof(e42, plain, ! [X12] : (sP35(X12) <=> (? [X17] : (~ p102(X17) & r1(X12, X17)) | ~ p602(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f43, plain, ! [X12] : (? [X16] : (~ p102(X16) & r1(X12, X16)) | ~ p502(X12) | ~ sP36(X12)), inference(usedef, [], [e43])).
fof(e43, plain, ! [X12] : (sP36(X12) <=> (? [X16] : (~ p102(X16) & r1(X12, X16)) | ~ p502(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f44, plain, ! [X12] : (? [X15] : (~ p102(X15) & r1(X12, X15)) | ~ p402(X12) | ~ sP37(X12)), inference(usedef, [], [e44])).
fof(e44, plain, ! [X12] : (sP37(X12) <=> (? [X15] : (~ p102(X15) & r1(X12, X15)) | ~ p402(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f45, plain, ! [X12] : (? [X14] : (~ p102(X14) & r1(X12, X14)) | ~ p302(X12) | ~ sP38(X12)), inference(usedef, [], [e45])).
fof(e45, plain, ! [X12] : (sP38(X12) <=> (? [X14] : (~ p102(X14) & r1(X12, X14)) | ~ p302(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f46, plain, ! [X12] : (? [X13] : (~ p102(X13) & r1(X12, X13)) | ~ p202(X12) | ~ sP39(X12)), inference(usedef, [], [e46])).
fof(e46, plain, ! [X12] : (sP39(X12) <=> (? [X13] : (~ p102(X13) & r1(X12, X13)) | ~ p202(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f6, plain, ? [X0] : (? [X1] : ((p101(X1) | ! [X2] : (p102(X2) | ~ r1(X1, X2)) | ! [X3] : (p103(X3) | ~ r1(X1, X3)) | ! [X4] : (p104(X4) | ~ r1(X1, X4)) | ! [X5] : (p105(X5) | ~ r1(X1, X5))) & (p201(X1) | p202(X1) | ! [X6] : (p203(X6) | ~ r1(X1, X6)) | ! [X7] : (p204(X7) | ~ r1(X1, X7)) | ! [X8] : (p205(X8) | ~ r1(X1, X8))) & (p301(X1) | p302(X1) | p303(X1) | ! [X9] : (p304(X9) | ~ r1(X1, X9)) | ! [X10] : (p305(X10) | ~ r1(X1, X10))) & (p401(X1) | p402(X1) | p403(X1) | p404(X1) | ! [X11] : (p405(X11) | ~ r1(X1, X11))) & (p501(X1) | p502(X1) | p503(X1) | p504(X1) | p505(X1)) & (p601(X1) | p602(X1) | p603(X1) | p604(X1) | p605(X1)) & r1(X0, X1)) & ! [X12] : (((~ p101(X12) | ~ p201(X12)) & (~ p101(X12) | ~ p301(X12)) & (~ p101(X12) | ~ p401(X12)) & (~ p101(X12) | ~ p501(X12)) & (~ p101(X12) | ~ p601(X12)) & (~ p201(X12) | ~ p301(X12)) & (~ p201(X12) | ~ p401(X12)) & (~ p201(X12) | ~ p501(X12)) & (~ p201(X12) | ~ p601(X12)) & (~ p301(X12) | ~ p401(X12)) & (~ p301(X12) | ~ p501(X12)) & (~ p301(X12) | ~ p601(X12)) & (~ p401(X12) | ~ p501(X12)) & (~ p401(X12) | ~ p601(X12)) & (~ p501(X12) | ~ p601(X12)) & (? [X13] : (~ p102(X13) & r1(X12, X13)) | ~ p202(X12)) & (? [X14] : (~ p102(X14) & r1(X12, X14)) | ~ p302(X12)) & (? [X15] : (~ p102(X15) & r1(X12, X15)) | ~ p402(X12)) & (? [X16] : (~ p102(X16) & r1(X12, X16)) | ~ p502(X12)) & (? [X17] : (~ p102(X17) & r1(X12, X17)) | ~ p602(X12)) & (~ p202(X12) | ~ p302(X12)) & (~ p202(X12) | ~ p402(X12)) & (~ p202(X12) | ~ p502(X12)) & (~ p202(X12) | ~ p602(X12)) & (~ p302(X12) | ~ p402(X12)) & (~ p302(X12) | ~ p502(X12)) & (~ p302(X12) | ~ p602(X12)) & (~ p402(X12) | ~ p502(X12)) & (~ p402(X12) | ~ p602(X12)) & (~ p502(X12) | ~ p602(X12)) & (? [X18] : (~ p103(X18) & r1(X12, X18)) | ? [X19] : (~ p203(X19) & r1(X12, X19))) & (? [X20] : (~ p103(X20) & r1(X12, X20)) | ~ p303(X12)) & (? [X21] : (~ p103(X21) & r1(X12, X21)) | ~ p403(X12)) & (? [X22] : (~ p103(X22) & r1(X12, X22)) | ~ p503(X12)) & (? [X23] : (~ p103(X23) & r1(X12, X23)) | ~ p603(X12)) & (? [X24] : (~ p203(X24) & r1(X12, X24)) | ~ p303(X12)) & (? [X25] : (~ p203(X25) & r1(X12, X25)) | ~ p403(X12)) & (? [X26] : (~ p203(X26) & r1(X12, X26)) | ~ p503(X12)) & (? [X27] : (~ p203(X27) & r1(X12, X27)) | ~ p603(X12)) & (~ p303(X12) | ~ p403(X12)) & (~ p303(X12) | ~ p503(X12)) & (~ p303(X12) | ~ p603(X12)) & (~ p403(X12) | ~ p503(X12)) & (~ p403(X12) | ~ p603(X12)) & (~ p503(X12) | ~ p603(X12)) & (? [X28] : (~ p104(X28) & r1(X12, X28)) | ? [X29] : (~ p204(X29) & r1(X12, X29))) & (? [X30] : (~ p104(X30) & r1(X12, X30)) | ? [X31] : (~ p304(X31) & r1(X12, X31))) & (? [X32] : (~ p104(X32) & r1(X12, X32)) | ~ p404(X12)) & (? [X33] : (~ p104(X33) & r1(X12, X33)) | ~ p504(X12)) & (? [X34] : (~ p104(X34) & r1(X12, X34)) | ~ p604(X12)) & (? [X35] : (~ p204(X35) & r1(X12, X35)) | ? [X36] : (~ p304(X36) & r1(X12, X36))) & (? [X37] : (~ p204(X37) & r1(X12, X37)) | ~ p404(X12)) & (? [X38] : (~ p204(X38) & r1(X12, X38)) | ~ p504(X12)) & (? [X39] : (~ p204(X39) & r1(X12, X39)) | ~ p604(X12)) & (? [X40] : (~ p304(X40) & r1(X12, X40)) | ~ p404(X12)) & (? [X41] : (~ p304(X41) & r1(X12, X41)) | ~ p504(X12)) & (? [X42] : (~ p304(X42) & r1(X12, X42)) | ~ p604(X12)) & (~ p404(X12) | ~ p504(X12)) & (~ p404(X12) | ~ p604(X12)) & (~ p504(X12) | ~ p604(X12)) & (? [X43] : (~ p105(X43) & r1(X12, X43)) | ? [X44] : (~ p205(X44) & r1(X12, X44))) & (? [X45] : (~ p105(X45) & r1(X12, X45)) | ? [X46] : (~ p305(X46) & r1(X12, X46))) & (? [X47] : (~ p105(X47) & r1(X12, X47)) | ? [X48] : (~ p405(X48) & r1(X12, X48))) & (? [X49] : (~ p105(X49) & r1(X12, X49)) | ~ p505(X12)) & (? [X50] : (~ p105(X50) & r1(X12, X50)) | ~ p605(X12)) & (? [X51] : (~ p205(X51) & r1(X12, X51)) | ? [X52] : (~ p305(X52) & r1(X12, X52))) & (? [X53] : (~ p205(X53) & r1(X12, X53)) | ? [X54] : (~ p405(X54) & r1(X12, X54))) & (? [X55] : (~ p205(X55) & r1(X12, X55)) | ~ p505(X12)) & (? [X56] : (~ p205(X56) & r1(X12, X56)) | ~ p605(X12)) & (? [X57] : (~ p305(X57) & r1(X12, X57)) | ? [X58] : (~ p405(X58) & r1(X12, X58))) & (? [X59] : (~ p305(X59) & r1(X12, X59)) | ~ p505(X12)) & (? [X60] : (~ p305(X60) & r1(X12, X60)) | ~ p605(X12)) & (? [X61] : (~ p405(X61) & r1(X12, X61)) | ~ p505(X12)) & (? [X62] : (~ p405(X62) & r1(X12, X62)) | ~ p605(X12)) & (~ p505(X12) | ~ p605(X12))) | ~ r1(X0, X12))), inference(flattening, [], [f5])).
fof(f5, plain, ? [X0] : (? [X1] : (((p101(X1) | ! [X2] : (p102(X2) | ~ r1(X1, X2)) | ! [X3] : (p103(X3) | ~ r1(X1, X3)) | ! [X4] : (p104(X4) | ~ r1(X1, X4)) | ! [X5] : (p105(X5) | ~ r1(X1, X5))) & (p201(X1) | p202(X1) | ! [X6] : (p203(X6) | ~ r1(X1, X6)) | ! [X7] : (p204(X7) | ~ r1(X1, X7)) | ! [X8] : (p205(X8) | ~ r1(X1, X8))) & (p301(X1) | p302(X1) | p303(X1) | ! [X9] : (p304(X9) | ~ r1(X1, X9)) | ! [X10] : (p305(X10) | ~ r1(X1, X10))) & (p401(X1) | p402(X1) | p403(X1) | p404(X1) | ! [X11] : (p405(X11) | ~ r1(X1, X11))) & (p501(X1) | p502(X1) | p503(X1) | p504(X1) | p505(X1)) & (p601(X1) | p602(X1) | p603(X1) | p604(X1) | p605(X1))) & r1(X0, X1)) & ! [X12] : (((~ p101(X12) | ~ p201(X12)) & (~ p101(X12) | ~ p301(X12)) & (~ p101(X12) | ~ p401(X12)) & (~ p101(X12) | ~ p501(X12)) & (~ p101(X12) | ~ p601(X12)) & (~ p201(X12) | ~ p301(X12)) & (~ p201(X12) | ~ p401(X12)) & (~ p201(X12) | ~ p501(X12)) & (~ p201(X12) | ~ p601(X12)) & (~ p301(X12) | ~ p401(X12)) & (~ p301(X12) | ~ p501(X12)) & (~ p301(X12) | ~ p601(X12)) & (~ p401(X12) | ~ p501(X12)) & (~ p401(X12) | ~ p601(X12)) & (~ p501(X12) | ~ p601(X12)) & (? [X13] : (~ p102(X13) & r1(X12, X13)) | ~ p202(X12)) & (? [X14] : (~ p102(X14) & r1(X12, X14)) | ~ p302(X12)) & (? [X15] : (~ p102(X15) & r1(X12, X15)) | ~ p402(X12)) & (? [X16] : (~ p102(X16) & r1(X12, X16)) | ~ p502(X12)) & (? [X17] : (~ p102(X17) & r1(X12, X17)) | ~ p602(X12)) & (~ p202(X12) | ~ p302(X12)) & (~ p202(X12) | ~ p402(X12)) & (~ p202(X12) | ~ p502(X12)) & (~ p202(X12) | ~ p602(X12)) & (~ p302(X12) | ~ p402(X12)) & (~ p302(X12) | ~ p502(X12)) & (~ p302(X12) | ~ p602(X12)) & (~ p402(X12) | ~ p502(X12)) & (~ p402(X12) | ~ p602(X12)) & (~ p502(X12) | ~ p602(X12)) & (? [X18] : (~ p103(X18) & r1(X12, X18)) | ? [X19] : (~ p203(X19) & r1(X12, X19))) & (? [X20] : (~ p103(X20) & r1(X12, X20)) | ~ p303(X12)) & (? [X21] : (~ p103(X21) & r1(X12, X21)) | ~ p403(X12)) & (? [X22] : (~ p103(X22) & r1(X12, X22)) | ~ p503(X12)) & (? [X23] : (~ p103(X23) & r1(X12, X23)) | ~ p603(X12)) & (? [X24] : (~ p203(X24) & r1(X12, X24)) | ~ p303(X12)) & (? [X25] : (~ p203(X25) & r1(X12, X25)) | ~ p403(X12)) & (? [X26] : (~ p203(X26) & r1(X12, X26)) | ~ p503(X12)) & (? [X27] : (~ p203(X27) & r1(X12, X27)) | ~ p603(X12)) & (~ p303(X12) | ~ p403(X12)) & (~ p303(X12) | ~ p503(X12)) & (~ p303(X12) | ~ p603(X12)) & (~ p403(X12) | ~ p503(X12)) & (~ p403(X12) | ~ p603(X12)) & (~ p503(X12) | ~ p603(X12)) & (? [X28] : (~ p104(X28) & r1(X12, X28)) | ? [X29] : (~ p204(X29) & r1(X12, X29))) & (? [X30] : (~ p104(X30) & r1(X12, X30)) | ? [X31] : (~ p304(X31) & r1(X12, X31))) & (? [X32] : (~ p104(X32) & r1(X12, X32)) | ~ p404(X12)) & (? [X33] : (~ p104(X33) & r1(X12, X33)) | ~ p504(X12)) & (? [X34] : (~ p104(X34) & r1(X12, X34)) | ~ p604(X12)) & (? [X35] : (~ p204(X35) & r1(X12, X35)) | ? [X36] : (~ p304(X36) & r1(X12, X36))) & (? [X37] : (~ p204(X37) & r1(X12, X37)) | ~ p404(X12)) & (? [X38] : (~ p204(X38) & r1(X12, X38)) | ~ p504(X12)) & (? [X39] : (~ p204(X39) & r1(X12, X39)) | ~ p604(X12)) & (? [X40] : (~ p304(X40) & r1(X12, X40)) | ~ p404(X12)) & (? [X41] : (~ p304(X41) & r1(X12, X41)) | ~ p504(X12)) & (? [X42] : (~ p304(X42) & r1(X12, X42)) | ~ p604(X12)) & (~ p404(X12) | ~ p504(X12)) & (~ p404(X12) | ~ p604(X12)) & (~ p504(X12) | ~ p604(X12)) & (? [X43] : (~ p105(X43) & r1(X12, X43)) | ? [X44] : (~ p205(X44) & r1(X12, X44))) & (? [X45] : (~ p105(X45) & r1(X12, X45)) | ? [X46] : (~ p305(X46) & r1(X12, X46))) & (? [X47] : (~ p105(X47) & r1(X12, X47)) | ? [X48] : (~ p405(X48) & r1(X12, X48))) & (? [X49] : (~ p105(X49) & r1(X12, X49)) | ~ p505(X12)) & (? [X50] : (~ p105(X50) & r1(X12, X50)) | ~ p605(X12)) & (? [X51] : (~ p205(X51) & r1(X12, X51)) | ? [X52] : (~ p305(X52) & r1(X12, X52))) & (? [X53] : (~ p205(X53) & r1(X12, X53)) | ? [X54] : (~ p405(X54) & r1(X12, X54))) & (? [X55] : (~ p205(X55) & r1(X12, X55)) | ~ p505(X12)) & (? [X56] : (~ p205(X56) & r1(X12, X56)) | ~ p605(X12)) & (? [X57] : (~ p305(X57) & r1(X12, X57)) | ? [X58] : (~ p405(X58) & r1(X12, X58))) & (? [X59] : (~ p305(X59) & r1(X12, X59)) | ~ p505(X12)) & (? [X60] : (~ p305(X60) & r1(X12, X60)) | ~ p605(X12)) & (? [X61] : (~ p405(X61) & r1(X12, X61)) | ~ p505(X12)) & (? [X62] : (~ p405(X62) & r1(X12, X62)) | ~ p605(X12)) & (~ p505(X12) | ~ p605(X12))) | ~ r1(X0, X12))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ? [X0] : ~ (! [X1] : (~ ((p101(X1) | ! [X2] : (p102(X2) | ~ r1(X1, X2)) | ! [X3] : (p103(X3) | ~ r1(X1, X3)) | ! [X4] : (p104(X4) | ~ r1(X1, X4)) | ! [X5] : (p105(X5) | ~ r1(X1, X5))) & (p201(X1) | p202(X1) | ! [X6] : (p203(X6) | ~ r1(X1, X6)) | ! [X7] : (p204(X7) | ~ r1(X1, X7)) | ! [X8] : (p205(X8) | ~ r1(X1, X8))) & (p301(X1) | p302(X1) | p303(X1) | ! [X9] : (p304(X9) | ~ r1(X1, X9)) | ! [X10] : (p305(X10) | ~ r1(X1, X10))) & (p401(X1) | p402(X1) | p403(X1) | p404(X1) | ! [X11] : (p405(X11) | ~ r1(X1, X11))) & (p501(X1) | p502(X1) | p503(X1) | p504(X1) | p505(X1)) & (p601(X1) | p602(X1) | p603(X1) | p604(X1) | p605(X1))) | ~ r1(X0, X1)) | ~ ! [X12] : (~ ((p101(X12) & p201(X12)) | (p101(X12) & p301(X12)) | (p101(X12) & p401(X12)) | (p101(X12) & p501(X12)) | (p101(X12) & p601(X12)) | (p201(X12) & p301(X12)) | (p201(X12) & p401(X12)) | (p201(X12) & p501(X12)) | (p201(X12) & p601(X12)) | (p301(X12) & p401(X12)) | (p301(X12) & p501(X12)) | (p301(X12) & p601(X12)) | (p401(X12) & p501(X12)) | (p401(X12) & p601(X12)) | (p501(X12) & p601(X12)) | (! [X13] : (p102(X13) | ~ r1(X12, X13)) & p202(X12)) | (! [X14] : (p102(X14) | ~ r1(X12, X14)) & p302(X12)) | (! [X15] : (p102(X15) | ~ r1(X12, X15)) & p402(X12)) | (! [X16] : (p102(X16) | ~ r1(X12, X16)) & p502(X12)) | (! [X17] : (p102(X17) | ~ r1(X12, X17)) & p602(X12)) | (p202(X12) & p302(X12)) | (p202(X12) & p402(X12)) | (p202(X12) & p502(X12)) | (p202(X12) & p602(X12)) | (p302(X12) & p402(X12)) | (p302(X12) & p502(X12)) | (p302(X12) & p602(X12)) | (p402(X12) & p502(X12)) | (p402(X12) & p602(X12)) | (p502(X12) & p602(X12)) | (! [X18] : (p103(X18) | ~ r1(X12, X18)) & ! [X19] : (p203(X19) | ~ r1(X12, X19))) | (! [X20] : (p103(X20) | ~ r1(X12, X20)) & p303(X12)) | (! [X21] : (p103(X21) | ~ r1(X12, X21)) & p403(X12)) | (! [X22] : (p103(X22) | ~ r1(X12, X22)) & p503(X12)) | (! [X23] : (p103(X23) | ~ r1(X12, X23)) & p603(X12)) | (! [X24] : (p203(X24) | ~ r1(X12, X24)) & p303(X12)) | (! [X25] : (p203(X25) | ~ r1(X12, X25)) & p403(X12)) | (! [X26] : (p203(X26) | ~ r1(X12, X26)) & p503(X12)) | (! [X27] : (p203(X27) | ~ r1(X12, X27)) & p603(X12)) | (p303(X12) & p403(X12)) | (p303(X12) & p503(X12)) | (p303(X12) & p603(X12)) | (p403(X12) & p503(X12)) | (p403(X12) & p603(X12)) | (p503(X12) & p603(X12)) | (! [X28] : (p104(X28) | ~ r1(X12, X28)) & ! [X29] : (p204(X29) | ~ r1(X12, X29))) | (! [X30] : (p104(X30) | ~ r1(X12, X30)) & ! [X31] : (p304(X31) | ~ r1(X12, X31))) | (! [X32] : (p104(X32) | ~ r1(X12, X32)) & p404(X12)) | (! [X33] : (p104(X33) | ~ r1(X12, X33)) & p504(X12)) | (! [X34] : (p104(X34) | ~ r1(X12, X34)) & p604(X12)) | (! [X35] : (p204(X35) | ~ r1(X12, X35)) & ! [X36] : (p304(X36) | ~ r1(X12, X36))) | (! [X37] : (p204(X37) | ~ r1(X12, X37)) & p404(X12)) | (! [X38] : (p204(X38) | ~ r1(X12, X38)) & p504(X12)) | (! [X39] : (p204(X39) | ~ r1(X12, X39)) & p604(X12)) | (! [X40] : (p304(X40) | ~ r1(X12, X40)) & p404(X12)) | (! [X41] : (p304(X41) | ~ r1(X12, X41)) & p504(X12)) | (! [X42] : (p304(X42) | ~ r1(X12, X42)) & p604(X12)) | (p404(X12) & p504(X12)) | (p404(X12) & p604(X12)) | (p504(X12) & p604(X12)) | (! [X43] : (p105(X43) | ~ r1(X12, X43)) & ! [X44] : (p205(X44) | ~ r1(X12, X44))) | (! [X45] : (p105(X45) | ~ r1(X12, X45)) & ! [X46] : (p305(X46) | ~ r1(X12, X46))) | (! [X47] : (p105(X47) | ~ r1(X12, X47)) & ! [X48] : (p405(X48) | ~ r1(X12, X48))) | (! [X49] : (p105(X49) | ~ r1(X12, X49)) & p505(X12)) | (! [X50] : (p105(X50) | ~ r1(X12, X50)) & p605(X12)) | (! [X51] : (p205(X51) | ~ r1(X12, X51)) & ! [X52] : (p305(X52) | ~ r1(X12, X52))) | (! [X53] : (p205(X53) | ~ r1(X12, X53)) & ! [X54] : (p405(X54) | ~ r1(X12, X54))) | (! [X55] : (p205(X55) | ~ r1(X12, X55)) & p505(X12)) | (! [X56] : (p205(X56) | ~ r1(X12, X56)) & p605(X12)) | (! [X57] : (p305(X57) | ~ r1(X12, X57)) & ! [X58] : (p405(X58) | ~ r1(X12, X58))) | (! [X59] : (p305(X59) | ~ r1(X12, X59)) & p505(X12)) | (! [X60] : (p305(X60) | ~ r1(X12, X60)) & p605(X12)) | (! [X61] : (p405(X61) | ~ r1(X12, X61)) & p505(X12)) | (! [X62] : (p405(X62) | ~ r1(X12, X62)) & p605(X12)) | (p505(X12) & p605(X12))) | ~ r1(X0, X12))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (! [X1] : (~ ((p101(X1) | ! [X2] : (p102(X2) | ~ r1(X1, X2)) | ! [X3] : (p103(X3) | ~ r1(X1, X3)) | ! [X4] : (p104(X4) | ~ r1(X1, X4)) | ! [X5] : (p105(X5) | ~ r1(X1, X5))) & (p201(X1) | p202(X1) | ! [X6] : (p203(X6) | ~ r1(X1, X6)) | ! [X7] : (p204(X7) | ~ r1(X1, X7)) | ! [X8] : (p205(X8) | ~ r1(X1, X8))) & (p301(X1) | p302(X1) | p303(X1) | ! [X9] : (p304(X9) | ~ r1(X1, X9)) | ! [X10] : (p305(X10) | ~ r1(X1, X10))) & (p401(X1) | p402(X1) | p403(X1) | p404(X1) | ! [X11] : (p405(X11) | ~ r1(X1, X11))) & (p501(X1) | p502(X1) | p503(X1) | p504(X1) | p505(X1)) & (p601(X1) | p602(X1) | p603(X1) | p604(X1) | p605(X1))) | ~ r1(X0, X1)) | ~ ! [X12] : (~ ((p101(X12) & p201(X12)) | (p101(X12) & p301(X12)) | (p101(X12) & p401(X12)) | (p101(X12) & p501(X12)) | (p101(X12) & p601(X12)) | (p201(X12) & p301(X12)) | (p201(X12) & p401(X12)) | (p201(X12) & p501(X12)) | (p201(X12) & p601(X12)) | (p301(X12) & p401(X12)) | (p301(X12) & p501(X12)) | (p301(X12) & p601(X12)) | (p401(X12) & p501(X12)) | (p401(X12) & p601(X12)) | (p501(X12) & p601(X12)) | (! [X13] : (p102(X13) | ~ r1(X12, X13)) & p202(X12)) | (! [X14] : (p102(X14) | ~ r1(X12, X14)) & p302(X12)) | (! [X15] : (p102(X15) | ~ r1(X12, X15)) & p402(X12)) | (! [X16] : (p102(X16) | ~ r1(X12, X16)) & p502(X12)) | (! [X17] : (p102(X17) | ~ r1(X12, X17)) & p602(X12)) | (p202(X12) & p302(X12)) | (p202(X12) & p402(X12)) | (p202(X12) & p502(X12)) | (p202(X12) & p602(X12)) | (p302(X12) & p402(X12)) | (p302(X12) & p502(X12)) | (p302(X12) & p602(X12)) | (p402(X12) & p502(X12)) | (p402(X12) & p602(X12)) | (p502(X12) & p602(X12)) | (! [X18] : (p103(X18) | ~ r1(X12, X18)) & ! [X19] : (p203(X19) | ~ r1(X12, X19))) | (! [X20] : (p103(X20) | ~ r1(X12, X20)) & p303(X12)) | (! [X21] : (p103(X21) | ~ r1(X12, X21)) & p403(X12)) | (! [X22] : (p103(X22) | ~ r1(X12, X22)) & p503(X12)) | (! [X23] : (p103(X23) | ~ r1(X12, X23)) & p603(X12)) | (! [X24] : (p203(X24) | ~ r1(X12, X24)) & p303(X12)) | (! [X25] : (p203(X25) | ~ r1(X12, X25)) & p403(X12)) | (! [X26] : (p203(X26) | ~ r1(X12, X26)) & p503(X12)) | (! [X27] : (p203(X27) | ~ r1(X12, X27)) & p603(X12)) | (p303(X12) & p403(X12)) | (p303(X12) & p503(X12)) | (p303(X12) & p603(X12)) | (p403(X12) & p503(X12)) | (p403(X12) & p603(X12)) | (p503(X12) & p603(X12)) | (! [X28] : (p104(X28) | ~ r1(X12, X28)) & ! [X29] : (p204(X29) | ~ r1(X12, X29))) | (! [X30] : (p104(X30) | ~ r1(X12, X30)) & ! [X31] : (p304(X31) | ~ r1(X12, X31))) | (! [X32] : (p104(X32) | ~ r1(X12, X32)) & p404(X12)) | (! [X33] : (p104(X33) | ~ r1(X12, X33)) & p504(X12)) | (! [X34] : (p104(X34) | ~ r1(X12, X34)) & p604(X12)) | (! [X35] : (p204(X35) | ~ r1(X12, X35)) & ! [X36] : (p304(X36) | ~ r1(X12, X36))) | (! [X37] : (p204(X37) | ~ r1(X12, X37)) & p404(X12)) | (! [X38] : (p204(X38) | ~ r1(X12, X38)) & p504(X12)) | (! [X39] : (p204(X39) | ~ r1(X12, X39)) & p604(X12)) | (! [X40] : (p304(X40) | ~ r1(X12, X40)) & p404(X12)) | (! [X41] : (p304(X41) | ~ r1(X12, X41)) & p504(X12)) | (! [X42] : (p304(X42) | ~ r1(X12, X42)) & p604(X12)) | (p404(X12) & p504(X12)) | (p404(X12) & p604(X12)) | (p504(X12) & p604(X12)) | (! [X43] : (p105(X43) | ~ r1(X12, X43)) & ! [X44] : (p205(X44) | ~ r1(X12, X44))) | (! [X45] : (p105(X45) | ~ r1(X12, X45)) & ! [X46] : (p305(X46) | ~ r1(X12, X46))) | (! [X47] : (p105(X47) | ~ r1(X12, X47)) & ! [X48] : (p405(X48) | ~ r1(X12, X48))) | (! [X49] : (p105(X49) | ~ r1(X12, X49)) & p505(X12)) | (! [X50] : (p105(X50) | ~ r1(X12, X50)) & p605(X12)) | (! [X51] : (p205(X51) | ~ r1(X12, X51)) & ! [X52] : (p305(X52) | ~ r1(X12, X52))) | (! [X53] : (p205(X53) | ~ r1(X12, X53)) & ! [X54] : (p405(X54) | ~ r1(X12, X54))) | (! [X55] : (p205(X55) | ~ r1(X12, X55)) & p505(X12)) | (! [X56] : (p205(X56) | ~ r1(X12, X56)) & p605(X12)) | (! [X57] : (p305(X57) | ~ r1(X12, X57)) & ! [X58] : (p405(X58) | ~ r1(X12, X58))) | (! [X59] : (p305(X59) | ~ r1(X12, X59)) & p505(X12)) | (! [X60] : (p305(X60) | ~ r1(X12, X60)) & p605(X12)) | (! [X61] : (p405(X61) | ~ r1(X12, X61)) & p505(X12)) | (! [X62] : (p405(X62) | ~ r1(X12, X62)) & p605(X12)) | (p505(X12) & p605(X12))) | ~ r1(X0, X12))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (! [X1] : (~ ((p101(X1) | ! [X0] : (p102(X0) | ~ r1(X1, X0)) | ! [X0] : (p103(X0) | ~ r1(X1, X0)) | ! [X0] : (p104(X0) | ~ r1(X1, X0)) | ! [X0] : (p105(X0) | ~ r1(X1, X0))) & (p201(X1) | p202(X1) | ! [X0] : (p203(X0) | ~ r1(X1, X0)) | ! [X0] : (p204(X0) | ~ r1(X1, X0)) | ! [X0] : (p205(X0) | ~ r1(X1, X0))) & (p301(X1) | p302(X1) | p303(X1) | ! [X0] : (p304(X0) | ~ r1(X1, X0)) | ! [X0] : (p305(X0) | ~ r1(X1, X0))) & (p401(X1) | p402(X1) | p403(X1) | p404(X1) | ! [X0] : (p405(X0) | ~ r1(X1, X0))) & (p501(X1) | p502(X1) | p503(X1) | p504(X1) | p505(X1)) & (p601(X1) | p602(X1) | p603(X1) | p604(X1) | p605(X1))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ((p101(X1) & p201(X1)) | (p101(X1) & p301(X1)) | (p101(X1) & p401(X1)) | (p101(X1) & p501(X1)) | (p101(X1) & p601(X1)) | (p201(X1) & p301(X1)) | (p201(X1) & p401(X1)) | (p201(X1) & p501(X1)) | (p201(X1) & p601(X1)) | (p301(X1) & p401(X1)) | (p301(X1) & p501(X1)) | (p301(X1) & p601(X1)) | (p401(X1) & p501(X1)) | (p401(X1) & p601(X1)) | (p501(X1) & p601(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p202(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p302(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p402(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p502(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p602(X1)) | (p202(X1) & p302(X1)) | (p202(X1) & p402(X1)) | (p202(X1) & p502(X1)) | (p202(X1) & p602(X1)) | (p302(X1) & p402(X1)) | (p302(X1) & p502(X1)) | (p302(X1) & p602(X1)) | (p402(X1) & p502(X1)) | (p402(X1) & p602(X1)) | (p502(X1) & p602(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & ! [X0] : (p203(X0) | ~ r1(X1, X0))) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p303(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p403(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p503(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p603(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p303(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p403(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p503(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p603(X1)) | (p303(X1) & p403(X1)) | (p303(X1) & p503(X1)) | (p303(X1) & p603(X1)) | (p403(X1) & p503(X1)) | (p403(X1) & p603(X1)) | (p503(X1) & p603(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & ! [X0] : (p204(X0) | ~ r1(X1, X0))) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & ! [X0] : (p304(X0) | ~ r1(X1, X0))) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p604(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & ! [X0] : (p304(X0) | ~ r1(X1, X0))) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p604(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p604(X1)) | (p404(X1) & p504(X1)) | (p404(X1) & p604(X1)) | (p504(X1) & p604(X1)) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p205(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p305(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & ! [X0] : (p305(X0) | ~ r1(X1, X0))) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p405(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p405(X0) | ~ r1(X1, X0)) & p605(X1)) | (p505(X1) & p605(X1))) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (! [X1] : (~ ((p101(X1) | ! [X0] : (p102(X0) | ~ r1(X1, X0)) | ! [X0] : (p103(X0) | ~ r1(X1, X0)) | ! [X0] : (p104(X0) | ~ r1(X1, X0)) | ! [X0] : (p105(X0) | ~ r1(X1, X0))) & (p201(X1) | p202(X1) | ! [X0] : (p203(X0) | ~ r1(X1, X0)) | ! [X0] : (p204(X0) | ~ r1(X1, X0)) | ! [X0] : (p205(X0) | ~ r1(X1, X0))) & (p301(X1) | p302(X1) | p303(X1) | ! [X0] : (p304(X0) | ~ r1(X1, X0)) | ! [X0] : (p305(X0) | ~ r1(X1, X0))) & (p401(X1) | p402(X1) | p403(X1) | p404(X1) | ! [X0] : (p405(X0) | ~ r1(X1, X0))) & (p501(X1) | p502(X1) | p503(X1) | p504(X1) | p505(X1)) & (p601(X1) | p602(X1) | p603(X1) | p604(X1) | p605(X1))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ((p101(X1) & p201(X1)) | (p101(X1) & p301(X1)) | (p101(X1) & p401(X1)) | (p101(X1) & p501(X1)) | (p101(X1) & p601(X1)) | (p201(X1) & p301(X1)) | (p201(X1) & p401(X1)) | (p201(X1) & p501(X1)) | (p201(X1) & p601(X1)) | (p301(X1) & p401(X1)) | (p301(X1) & p501(X1)) | (p301(X1) & p601(X1)) | (p401(X1) & p501(X1)) | (p401(X1) & p601(X1)) | (p501(X1) & p601(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p202(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p302(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p402(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p502(X1)) | (! [X0] : (p102(X0) | ~ r1(X1, X0)) & p602(X1)) | (p202(X1) & p302(X1)) | (p202(X1) & p402(X1)) | (p202(X1) & p502(X1)) | (p202(X1) & p602(X1)) | (p302(X1) & p402(X1)) | (p302(X1) & p502(X1)) | (p302(X1) & p602(X1)) | (p402(X1) & p502(X1)) | (p402(X1) & p602(X1)) | (p502(X1) & p602(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & ! [X0] : (p203(X0) | ~ r1(X1, X0))) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p303(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p403(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p503(X1)) | (! [X0] : (p103(X0) | ~ r1(X1, X0)) & p603(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p303(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p403(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p503(X1)) | (! [X0] : (p203(X0) | ~ r1(X1, X0)) & p603(X1)) | (p303(X1) & p403(X1)) | (p303(X1) & p503(X1)) | (p303(X1) & p603(X1)) | (p403(X1) & p503(X1)) | (p403(X1) & p603(X1)) | (p503(X1) & p603(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & ! [X0] : (p204(X0) | ~ r1(X1, X0))) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & ! [X0] : (p304(X0) | ~ r1(X1, X0))) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p104(X0) | ~ r1(X1, X0)) & p604(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & ! [X0] : (p304(X0) | ~ r1(X1, X0))) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p204(X0) | ~ r1(X1, X0)) & p604(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p404(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p504(X1)) | (! [X0] : (p304(X0) | ~ r1(X1, X0)) & p604(X1)) | (p404(X1) & p504(X1)) | (p404(X1) & p604(X1)) | (p504(X1) & p604(X1)) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p205(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p305(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p105(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & ! [X0] : (p305(X0) | ~ r1(X1, X0))) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p205(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & ! [X0] : (p405(X0) | ~ r1(X1, X0))) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p305(X0) | ~ r1(X1, X0)) & p605(X1)) | (! [X0] : (p405(X0) | ~ r1(X1, X0)) & p505(X1)) | (! [X0] : (p405(X0) | ~ r1(X1, X0)) & p605(X1)) | (p505(X1) & p605(X1))) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL648+1.005.p', main)).
fof(f399, plain, ! [X12] : (~ r1(sK91, X12) | sP40(X12)), inference(cnf_transformation, [], [f223])).
fof(f6576, plain, (~ sP33(sK92) | (~ spl93_3 | ~ spl93_18)), inference(subsumption_resolution, [], [f6575, f471])).
fof(f471, plain, (p403(sK92) | ~ spl93_18), inference(avatar_component_clause, [], [f469])).
fof(f469, plain, (spl93_18 <=> p403(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_18])])).
fof(f6575, plain, (~ p403(sK92) | ~ sP33(sK92) | (~ spl93_3 | ~ spl93_18)), inference(resolution, [], [f6485, f312])).
fof(f312, plain, ! [X0] : (~ p103(sK47(X0)) | ~ p403(X0) | ~ sP33(X0)), inference(cnf_transformation, [], [f78])).
fof(f78, plain, ! [X0] : ((~ p103(sK47(X0)) & r1(X0, sK47(X0))) | ~ p403(X0) | ~ sP33(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK47])], [f76, f77])).
fof(f77, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) => (~ p103(sK47(X0)) & r1(X0, sK47(X0)))), introduced(choice_axiom, [])).
fof(f76, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) | ~ p403(X0) | ~ sP33(X0)), inference(rectify, [], [f75])).
fof(f75, plain, ! [X12] : (? [X21] : (~ p103(X21) & r1(X12, X21)) | ~ p403(X12) | ~ sP33(X12)), inference(nnf_transformation, [], [f40])).
fof(f6485, plain, (p103(sK47(sK92)) | (~ spl93_3 | ~ spl93_18)), inference(subsumption_resolution, [], [f6484, f556])).
fof(f6484, plain, (p103(sK47(sK92)) | ~ sP33(sK92) | (~ spl93_3 | ~ spl93_18)), inference(subsumption_resolution, [], [f6430, f471])).
fof(f6430, plain, (p103(sK47(sK92)) | ~ p403(sK92) | ~ sP33(sK92) | ~ spl93_3), inference(resolution, [], [f415, f311])).
fof(f311, plain, ! [X0] : (r1(X0, sK47(X0)) | ~ p403(X0) | ~ sP33(X0)), inference(cnf_transformation, [], [f78])).
fof(f415, plain, (! [X3] : (~ r1(sK92, X3) | p103(X3)) | ~ spl93_3), inference(avatar_component_clause, [], [f414])).
fof(f414, plain, (spl93_3 <=> ! [X3] : (p103(X3) | ~ r1(sK92, X3))), introduced(avatar_definition, [new_symbols(naming, [spl93_3])])).
fof(f6423, plain, (~ spl93_8 | ~ spl93_18), inference(avatar_contradiction_clause, [], [f6422])).
fof(f6422, plain, ($false | (~ spl93_8 | ~ spl93_18)), inference(subsumption_resolution, [], [f6421, f552])).
fof(f552, plain, sP29(sK92), inference(resolution, [], [f523, f262])).
fof(f262, plain, ! [X0] : (~ sP40(X0) | sP29(X0)), inference(cnf_transformation, [], [f50])).
fof(f6421, plain, (~ sP29(sK92) | (~ spl93_8 | ~ spl93_18)), inference(subsumption_resolution, [], [f6420, f471])).
fof(f6420, plain, (~ p403(sK92) | ~ sP29(sK92) | (~ spl93_8 | ~ spl93_18)), inference(resolution, [], [f6403, f320])).
fof(f320, plain, ! [X0] : (~ p203(sK51(X0)) | ~ p403(X0) | ~ sP29(X0)), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ! [X0] : ((~ p203(sK51(X0)) & r1(X0, sK51(X0))) | ~ p403(X0) | ~ sP29(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK51])], [f92, f93])).
fof(f93, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) => (~ p203(sK51(X0)) & r1(X0, sK51(X0)))), introduced(choice_axiom, [])).
fof(f92, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) | ~ p403(X0) | ~ sP29(X0)), inference(rectify, [], [f91])).
fof(f91, plain, ! [X12] : (? [X25] : (~ p203(X25) & r1(X12, X25)) | ~ p403(X12) | ~ sP29(X12)), inference(nnf_transformation, [], [f36])).
fof(f6403, plain, (p203(sK51(sK92)) | (~ spl93_8 | ~ spl93_18)), inference(subsumption_resolution, [], [f6402, f552])).
fof(f6402, plain, (p203(sK51(sK92)) | ~ sP29(sK92) | (~ spl93_8 | ~ spl93_18)), inference(subsumption_resolution, [], [f6350, f471])).
fof(f6350, plain, (p203(sK51(sK92)) | ~ p403(sK92) | ~ sP29(sK92) | ~ spl93_8), inference(resolution, [], [f432, f319])).
fof(f319, plain, ! [X0] : (r1(X0, sK51(X0)) | ~ p403(X0) | ~ sP29(X0)), inference(cnf_transformation, [], [f94])).
fof(f432, plain, (! [X6] : (~ r1(sK92, X6) | p203(X6)) | ~ spl93_8), inference(avatar_component_clause, [], [f431])).
fof(f431, plain, (spl93_8 <=> ! [X6] : (p203(X6) | ~ r1(sK92, X6))), introduced(avatar_definition, [new_symbols(naming, [spl93_8])])).
fof(f6261, plain, (~ spl93_17 | ~ spl93_22), inference(avatar_split_clause, [], [f573, f486, f465])).
fof(f465, plain, (spl93_17 <=> p404(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_17])])).
fof(f486, plain, (spl93_22 <=> p504(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_22])])).
fof(f573, plain, (~ p504(sK92) | ~ p404(sK92)), inference(resolution, [], [f241, f523])).
fof(f241, plain, ! [X0] : (~ sP40(X0) | ~ p504(X0) | ~ p404(X0)), inference(cnf_transformation, [], [f50])).
fof(f6260, plain, (~ spl93_21 | ~ spl93_452), inference(avatar_contradiction_clause, [], [f6259])).
fof(f6259, plain, ($false | (~ spl93_21 | ~ spl93_452)), inference(subsumption_resolution, [], [f6258, f527])).
fof(f527, plain, sP13(sK92), inference(resolution, [], [f523, f228])).
fof(f228, plain, ! [X0] : (~ sP40(X0) | sP13(X0)), inference(cnf_transformation, [], [f50])).
fof(f6258, plain, (~ sP13(sK92) | (~ spl93_21 | ~ spl93_452)), inference(subsumption_resolution, [], [f6257, f484])).
fof(f484, plain, (p505(sK92) | ~ spl93_21), inference(avatar_component_clause, [], [f482])).
fof(f482, plain, (spl93_21 <=> p505(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_21])])).
fof(f6257, plain, (~ p505(sK92) | ~ sP13(sK92) | ~ spl93_452), inference(resolution, [], [f6237, f352])).
fof(f352, plain, ! [X0] : (~ p305(sK67(X0)) | ~ p505(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ! [X0] : ((~ p305(sK67(X0)) & r1(X0, sK67(X0))) | ~ p505(X0) | ~ sP13(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK67])], [f156, f157])).
fof(f157, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) => (~ p305(sK67(X0)) & r1(X0, sK67(X0)))), introduced(choice_axiom, [])).
fof(f156, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) | ~ p505(X0) | ~ sP13(X0)), inference(rectify, [], [f155])).
fof(f155, plain, ! [X12] : (? [X59] : (~ p305(X59) & r1(X12, X59)) | ~ p505(X12) | ~ sP13(X12)), inference(nnf_transformation, [], [f20])).
fof(f6237, plain, (p305(sK67(sK92)) | ~ spl93_452), inference(avatar_component_clause, [], [f6235])).
fof(f6235, plain, (spl93_452 <=> p305(sK67(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_452])])).
fof(f6238, plain, (~ spl93_21 | spl93_452 | ~ spl93_11), inference(avatar_split_clause, [], [f6233, f443, f6235, f482])).
fof(f443, plain, (spl93_11 <=> ! [X10] : (p305(X10) | ~ r1(sK92, X10))), introduced(avatar_definition, [new_symbols(naming, [spl93_11])])).
fof(f6233, plain, (p305(sK67(sK92)) | ~ p505(sK92) | ~ spl93_11), inference(subsumption_resolution, [], [f6083, f527])).
fof(f6083, plain, (p305(sK67(sK92)) | ~ p505(sK92) | ~ sP13(sK92) | ~ spl93_11), inference(resolution, [], [f444, f351])).
fof(f351, plain, ! [X0] : (r1(X0, sK67(X0)) | ~ p505(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f158])).
fof(f444, plain, (! [X10] : (~ r1(sK92, X10) | p305(X10)) | ~ spl93_11), inference(avatar_component_clause, [], [f443])).
fof(f6196, plain, (~ spl93_26 | ~ spl93_127), inference(avatar_contradiction_clause, [], [f6195])).
fof(f6195, plain, ($false | (~ spl93_26 | ~ spl93_127)), inference(subsumption_resolution, [], [f6194, f526])).
fof(f526, plain, sP12(sK92), inference(resolution, [], [f523, f227])).
fof(f227, plain, ! [X0] : (~ sP40(X0) | sP12(X0)), inference(cnf_transformation, [], [f50])).
fof(f6194, plain, (~ sP12(sK92) | (~ spl93_26 | ~ spl93_127)), inference(subsumption_resolution, [], [f6193, f505])).
fof(f505, plain, (p605(sK92) | ~ spl93_26), inference(avatar_component_clause, [], [f503])).
fof(f503, plain, (spl93_26 <=> p605(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_26])])).
fof(f6193, plain, (~ p605(sK92) | ~ sP12(sK92) | ~ spl93_127), inference(resolution, [], [f1319, f354])).
fof(f354, plain, ! [X0] : (~ p305(sK68(X0)) | ~ p605(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ! [X0] : ((~ p305(sK68(X0)) & r1(X0, sK68(X0))) | ~ p605(X0) | ~ sP12(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK68])], [f160, f161])).
fof(f161, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) => (~ p305(sK68(X0)) & r1(X0, sK68(X0)))), introduced(choice_axiom, [])).
fof(f160, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) | ~ p605(X0) | ~ sP12(X0)), inference(rectify, [], [f159])).
fof(f159, plain, ! [X12] : (? [X60] : (~ p305(X60) & r1(X12, X60)) | ~ p605(X12) | ~ sP12(X12)), inference(nnf_transformation, [], [f19])).
fof(f1319, plain, (p305(sK68(sK92)) | ~ spl93_127), inference(avatar_component_clause, [], [f1317])).
fof(f1317, plain, (spl93_127 <=> p305(sK68(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_127])])).
fof(f6128, plain, (~ spl93_26 | spl93_127 | ~ spl93_11), inference(avatar_split_clause, [], [f6127, f443, f1317, f503])).
fof(f6127, plain, (p305(sK68(sK92)) | ~ p605(sK92) | ~ spl93_11), inference(subsumption_resolution, [], [f6084, f526])).
fof(f6084, plain, (p305(sK68(sK92)) | ~ p605(sK92) | ~ sP12(sK92) | ~ spl93_11), inference(resolution, [], [f444, f353])).
fof(f353, plain, ! [X0] : (r1(X0, sK68(X0)) | ~ p605(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f162])).
fof(f6124, plain, (~ spl93_8 | ~ spl93_23), inference(avatar_contradiction_clause, [], [f6123])).
fof(f6123, plain, ($false | (~ spl93_8 | ~ spl93_23)), inference(subsumption_resolution, [], [f6122, f551])).
fof(f551, plain, sP28(sK92), inference(resolution, [], [f523, f261])).
fof(f261, plain, ! [X0] : (~ sP40(X0) | sP28(X0)), inference(cnf_transformation, [], [f50])).
fof(f6122, plain, (~ sP28(sK92) | (~ spl93_8 | ~ spl93_23)), inference(subsumption_resolution, [], [f6121, f492])).
fof(f492, plain, (p503(sK92) | ~ spl93_23), inference(avatar_component_clause, [], [f490])).
fof(f490, plain, (spl93_23 <=> p503(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_23])])).
fof(f6121, plain, (~ p503(sK92) | ~ sP28(sK92) | (~ spl93_8 | ~ spl93_23)), inference(resolution, [], [f6048, f322])).
fof(f322, plain, ! [X0] : (~ p203(sK52(X0)) | ~ p503(X0) | ~ sP28(X0)), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ! [X0] : ((~ p203(sK52(X0)) & r1(X0, sK52(X0))) | ~ p503(X0) | ~ sP28(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK52])], [f96, f97])).
fof(f97, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) => (~ p203(sK52(X0)) & r1(X0, sK52(X0)))), introduced(choice_axiom, [])).
fof(f96, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) | ~ p503(X0) | ~ sP28(X0)), inference(rectify, [], [f95])).
fof(f95, plain, ! [X12] : (? [X26] : (~ p203(X26) & r1(X12, X26)) | ~ p503(X12) | ~ sP28(X12)), inference(nnf_transformation, [], [f35])).
fof(f6048, plain, (p203(sK52(sK92)) | (~ spl93_8 | ~ spl93_23)), inference(subsumption_resolution, [], [f6047, f551])).
fof(f6047, plain, (p203(sK52(sK92)) | ~ sP28(sK92) | (~ spl93_8 | ~ spl93_23)), inference(subsumption_resolution, [], [f5994, f492])).
fof(f5994, plain, (p203(sK52(sK92)) | ~ p503(sK92) | ~ sP28(sK92) | ~ spl93_8), inference(resolution, [], [f432, f321])).
fof(f321, plain, ! [X0] : (r1(X0, sK52(X0)) | ~ p503(X0) | ~ sP28(X0)), inference(cnf_transformation, [], [f98])).
fof(f5982, plain, (~ spl93_5 | ~ spl93_20), inference(avatar_split_clause, [], [f616, f477, f420])).
fof(f420, plain, (spl93_5 <=> p101(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_5])])).
fof(f477, plain, (spl93_20 <=> p401(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_20])])).
fof(f616, plain, (~ p401(sK92) | ~ p101(sK92)), inference(resolution, [], [f296, f523])).
fof(f296, plain, ! [X0] : (~ sP40(X0) | ~ p401(X0) | ~ p101(X0)), inference(cnf_transformation, [], [f50])).
fof(f5981, plain, (~ spl93_15 | ~ spl93_20), inference(avatar_split_clause, [], [f607, f477, f457])).
fof(f457, plain, (spl93_15 <=> p301(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_15])])).
fof(f607, plain, (~ p401(sK92) | ~ p301(sK92)), inference(resolution, [], [f289, f523])).
fof(f289, plain, ! [X0] : (~ sP40(X0) | ~ p401(X0) | ~ p301(X0)), inference(cnf_transformation, [], [f50])).
fof(f5980, plain, (~ spl93_5 | ~ spl93_15), inference(avatar_split_clause, [], [f617, f457, f420])).
fof(f617, plain, (~ p301(sK92) | ~ p101(sK92)), inference(resolution, [], [f297, f523])).
fof(f297, plain, ! [X0] : (~ sP40(X0) | ~ p301(X0) | ~ p101(X0)), inference(cnf_transformation, [], [f50])).
fof(f5973, plain, (~ spl93_14 | ~ spl93_19), inference(avatar_split_clause, [], [f593, f473, f453])).
fof(f453, plain, (spl93_14 <=> p302(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_14])])).
fof(f473, plain, (spl93_19 <=> p402(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_19])])).
fof(f593, plain, (~ p402(sK92) | ~ p302(sK92)), inference(resolution, [], [f274, f523])).
fof(f274, plain, ! [X0] : (~ sP40(X0) | ~ p402(X0) | ~ p302(X0)), inference(cnf_transformation, [], [f50])).
fof(f5946, plain, (~ spl93_17 | ~ spl93_27), inference(avatar_split_clause, [], [f5945, f507, f465])).
fof(f507, plain, (spl93_27 <=> p604(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_27])])).
fof(f5945, plain, (~ p404(sK92) | ~ spl93_27), inference(subsumption_resolution, [], [f572, f509])).
fof(f509, plain, (p604(sK92) | ~ spl93_27), inference(avatar_component_clause, [], [f507])).
fof(f572, plain, (~ p604(sK92) | ~ p404(sK92)), inference(resolution, [], [f240, f523])).
fof(f240, plain, ! [X0] : (~ sP40(X0) | ~ p604(X0) | ~ p404(X0)), inference(cnf_transformation, [], [f50])).
fof(f5932, plain, (~ spl93_4 | ~ spl93_14), inference(avatar_contradiction_clause, [], [f5931])).
fof(f5931, plain, ($false | (~ spl93_4 | ~ spl93_14)), inference(subsumption_resolution, [], [f5930, f562])).
fof(f562, plain, sP38(sK92), inference(resolution, [], [f523, f282])).
fof(f282, plain, ! [X0] : (~ sP40(X0) | sP38(X0)), inference(cnf_transformation, [], [f50])).
fof(f5930, plain, (~ sP38(sK92) | (~ spl93_4 | ~ spl93_14)), inference(subsumption_resolution, [], [f5929, f455])).
fof(f455, plain, (p302(sK92) | ~ spl93_14), inference(avatar_component_clause, [], [f453])).
fof(f5929, plain, (~ p302(sK92) | ~ sP38(sK92) | (~ spl93_4 | ~ spl93_14)), inference(resolution, [], [f5842, f302])).
fof(f302, plain, ! [X0] : (~ p102(sK42(X0)) | ~ p302(X0) | ~ sP38(X0)), inference(cnf_transformation, [], [f58])).
fof(f58, plain, ! [X0] : ((~ p102(sK42(X0)) & r1(X0, sK42(X0))) | ~ p302(X0) | ~ sP38(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42])], [f56, f57])).
fof(f57, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) => (~ p102(sK42(X0)) & r1(X0, sK42(X0)))), introduced(choice_axiom, [])).
fof(f56, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) | ~ p302(X0) | ~ sP38(X0)), inference(rectify, [], [f55])).
fof(f55, plain, ! [X12] : (? [X14] : (~ p102(X14) & r1(X12, X14)) | ~ p302(X12) | ~ sP38(X12)), inference(nnf_transformation, [], [f45])).
fof(f5842, plain, (p102(sK42(sK92)) | (~ spl93_4 | ~ spl93_14)), inference(subsumption_resolution, [], [f5841, f562])).
fof(f5841, plain, (p102(sK42(sK92)) | ~ sP38(sK92) | (~ spl93_4 | ~ spl93_14)), inference(subsumption_resolution, [], [f5782, f455])).
fof(f5782, plain, (p102(sK42(sK92)) | ~ p302(sK92) | ~ sP38(sK92) | ~ spl93_4), inference(resolution, [], [f418, f301])).
fof(f301, plain, ! [X0] : (r1(X0, sK42(X0)) | ~ p302(X0) | ~ sP38(X0)), inference(cnf_transformation, [], [f58])).
fof(f418, plain, (! [X2] : (~ r1(sK92, X2) | p102(X2)) | ~ spl93_4), inference(avatar_component_clause, [], [f417])).
fof(f417, plain, (spl93_4 <=> ! [X2] : (p102(X2) | ~ r1(sK92, X2))), introduced(avatar_definition, [new_symbols(naming, [spl93_4])])).
fof(f5780, plain, (spl93_289 | spl93_235 | ~ spl93_11), inference(avatar_split_clause, [], [f5779, f443, f1984, f2804])).
fof(f2804, plain, (spl93_289 <=> r1(sK92, sK85(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_289])])).
fof(f1984, plain, (spl93_235 <=> p305(sK86(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_235])])).
fof(f5779, plain, (p305(sK86(sK92)) | r1(sK92, sK85(sK92)) | ~ spl93_11), inference(subsumption_resolution, [], [f5661, f532])).
fof(f532, plain, sP2(sK92), inference(resolution, [], [f523, f233])).
fof(f233, plain, ! [X0] : (~ sP40(X0) | sP2(X0)), inference(cnf_transformation, [], [f50])).
fof(f5661, plain, (p305(sK86(sK92)) | r1(sK92, sK85(sK92)) | ~ sP2(sK92) | ~ spl93_11), inference(resolution, [], [f444, f387])).
fof(f387, plain, ! [X0] : (r1(X0, sK86(X0)) | r1(X0, sK85(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f210])).
fof(f210, plain, ! [X0] : ((~ p205(sK85(X0)) & r1(X0, sK85(X0))) | (~ p305(sK86(X0)) & r1(X0, sK86(X0))) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK85, sK86])], [f207, f209, f208])).
fof(f208, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) => (~ p205(sK85(X0)) & r1(X0, sK85(X0)))), introduced(choice_axiom, [])).
fof(f209, plain, ! [X0] : (? [X2] : (~ p305(X2) & r1(X0, X2)) => (~ p305(sK86(X0)) & r1(X0, sK86(X0)))), introduced(choice_axiom, [])).
fof(f207, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) | ? [X2] : (~ p305(X2) & r1(X0, X2)) | ~ sP2(X0)), inference(rectify, [], [f206])).
fof(f206, plain, ! [X12] : (? [X51] : (~ p205(X51) & r1(X12, X51)) | ? [X52] : (~ p305(X52) & r1(X12, X52)) | ~ sP2(X12)), inference(nnf_transformation, [], [f9])).
fof(f5778, plain, (~ spl93_2 | ~ spl93_17), inference(avatar_contradiction_clause, [], [f5777])).
fof(f5777, plain, ($false | (~ spl93_2 | ~ spl93_17)), inference(subsumption_resolution, [], [f5776, f547])).
fof(f547, plain, sP26(sK92), inference(resolution, [], [f523, f251])).
fof(f251, plain, ! [X0] : (~ sP40(X0) | sP26(X0)), inference(cnf_transformation, [], [f50])).
fof(f5776, plain, (~ sP26(sK92) | (~ spl93_2 | ~ spl93_17)), inference(subsumption_resolution, [], [f5775, f467])).
fof(f467, plain, (p404(sK92) | ~ spl93_17), inference(avatar_component_clause, [], [f465])).
fof(f5775, plain, (~ p404(sK92) | ~ sP26(sK92) | (~ spl93_2 | ~ spl93_17)), inference(resolution, [], [f5768, f326])).
fof(f326, plain, ! [X0] : (~ p104(sK54(X0)) | ~ p404(X0) | ~ sP26(X0)), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ! [X0] : ((~ p104(sK54(X0)) & r1(X0, sK54(X0))) | ~ p404(X0) | ~ sP26(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54])], [f104, f105])).
fof(f105, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) => (~ p104(sK54(X0)) & r1(X0, sK54(X0)))), introduced(choice_axiom, [])).
fof(f104, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) | ~ p404(X0) | ~ sP26(X0)), inference(rectify, [], [f103])).
fof(f103, plain, ! [X12] : (? [X32] : (~ p104(X32) & r1(X12, X32)) | ~ p404(X12) | ~ sP26(X12)), inference(nnf_transformation, [], [f33])).
fof(f5768, plain, (p104(sK54(sK92)) | (~ spl93_2 | ~ spl93_17)), inference(subsumption_resolution, [], [f5767, f547])).
fof(f5767, plain, (p104(sK54(sK92)) | ~ sP26(sK92) | (~ spl93_2 | ~ spl93_17)), inference(subsumption_resolution, [], [f5718, f467])).
fof(f5718, plain, (p104(sK54(sK92)) | ~ p404(sK92) | ~ sP26(sK92) | ~ spl93_2), inference(resolution, [], [f412, f325])).
fof(f325, plain, ! [X0] : (r1(X0, sK54(X0)) | ~ p404(X0) | ~ sP26(X0)), inference(cnf_transformation, [], [f106])).
fof(f412, plain, (! [X4] : (~ r1(sK92, X4) | p104(X4)) | ~ spl93_2), inference(avatar_component_clause, [], [f411])).
fof(f411, plain, (spl93_2 <=> ! [X4] : (p104(X4) | ~ r1(sK92, X4))), introduced(avatar_definition, [new_symbols(naming, [spl93_2])])).
fof(f5702, plain, (spl93_286 | spl93_215 | ~ spl93_11), inference(avatar_split_clause, [], [f5701, f443, f1855, f2784])).
fof(f2784, plain, (spl93_286 <=> r1(sK92, sK81(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_286])])).
fof(f1855, plain, (spl93_215 <=> p305(sK82(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_215])])).
fof(f5701, plain, (p305(sK82(sK92)) | r1(sK92, sK81(sK92)) | ~ spl93_11), inference(subsumption_resolution, [], [f5655, f536])).
fof(f536, plain, sP4(sK92), inference(resolution, [], [f523, f237])).
fof(f237, plain, ! [X0] : (~ sP40(X0) | sP4(X0)), inference(cnf_transformation, [], [f50])).
fof(f5655, plain, (p305(sK82(sK92)) | r1(sK92, sK81(sK92)) | ~ sP4(sK92) | ~ spl93_11), inference(resolution, [], [f444, f379])).
fof(f379, plain, ! [X0] : (r1(X0, sK82(X0)) | r1(X0, sK81(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f200])).
fof(f200, plain, ! [X0] : ((~ p105(sK81(X0)) & r1(X0, sK81(X0))) | (~ p305(sK82(X0)) & r1(X0, sK82(X0))) | ~ sP4(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK81, sK82])], [f197, f199, f198])).
fof(f198, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) => (~ p105(sK81(X0)) & r1(X0, sK81(X0)))), introduced(choice_axiom, [])).
fof(f199, plain, ! [X0] : (? [X2] : (~ p305(X2) & r1(X0, X2)) => (~ p305(sK82(X0)) & r1(X0, sK82(X0)))), introduced(choice_axiom, [])).
fof(f197, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) | ? [X2] : (~ p305(X2) & r1(X0, X2)) | ~ sP4(X0)), inference(rectify, [], [f196])).
fof(f196, plain, ! [X12] : (? [X45] : (~ p105(X45) & r1(X12, X45)) | ? [X46] : (~ p305(X46) & r1(X12, X46)) | ~ sP4(X12)), inference(nnf_transformation, [], [f11])).
fof(f5700, plain, (~ spl93_220 | spl93_215 | ~ spl93_11), inference(avatar_split_clause, [], [f5699, f443, f1855, f1888])).
fof(f1888, plain, (spl93_220 <=> p105(sK81(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_220])])).
fof(f5699, plain, (p305(sK82(sK92)) | ~ p105(sK81(sK92)) | ~ spl93_11), inference(subsumption_resolution, [], [f5656, f536])).
fof(f5656, plain, (p305(sK82(sK92)) | ~ p105(sK81(sK92)) | ~ sP4(sK92) | ~ spl93_11), inference(resolution, [], [f444, f381])).
fof(f381, plain, ! [X0] : (r1(X0, sK82(X0)) | ~ p105(sK81(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f200])).
fof(f5698, plain, (~ spl93_215 | ~ spl93_220), inference(avatar_contradiction_clause, [], [f5697])).
fof(f5697, plain, ($false | (~ spl93_215 | ~ spl93_220)), inference(subsumption_resolution, [], [f5696, f536])).
fof(f5696, plain, (~ sP4(sK92) | (~ spl93_215 | ~ spl93_220)), inference(subsumption_resolution, [], [f5695, f1856])).
fof(f1856, plain, (p305(sK82(sK92)) | ~ spl93_215), inference(avatar_component_clause, [], [f1855])).
fof(f5695, plain, (~ p305(sK82(sK92)) | ~ sP4(sK92) | ~ spl93_220), inference(resolution, [], [f1889, f382])).
fof(f382, plain, ! [X0] : (~ p105(sK81(X0)) | ~ p305(sK82(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f200])).
fof(f1889, plain, (p105(sK81(sK92)) | ~ spl93_220), inference(avatar_component_clause, [], [f1888])).
fof(f5694, plain, (spl93_220 | ~ spl93_1 | ~ spl93_215), inference(avatar_split_clause, [], [f5693, f1855, f408, f1888])).
fof(f408, plain, (spl93_1 <=> ! [X5] : (p105(X5) | ~ r1(sK92, X5))), introduced(avatar_definition, [new_symbols(naming, [spl93_1])])).
fof(f5693, plain, (p105(sK81(sK92)) | (~ spl93_1 | ~ spl93_215)), inference(subsumption_resolution, [], [f5692, f536])).
fof(f5692, plain, (p105(sK81(sK92)) | ~ sP4(sK92) | (~ spl93_1 | ~ spl93_215)), inference(subsumption_resolution, [], [f5317, f1856])).
fof(f5317, plain, (p105(sK81(sK92)) | ~ p305(sK82(sK92)) | ~ sP4(sK92) | ~ spl93_1), inference(resolution, [], [f409, f380])).
fof(f380, plain, ! [X0] : (r1(X0, sK81(X0)) | ~ p305(sK82(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f200])).
fof(f409, plain, (! [X5] : (~ r1(sK92, X5) | p105(X5)) | ~ spl93_1), inference(avatar_component_clause, [], [f408])).
fof(f5594, plain, (~ spl93_235 | ~ spl93_237), inference(avatar_split_clause, [], [f5593, f1996, f1984])).
fof(f1996, plain, (spl93_237 <=> p205(sK85(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_237])])).
fof(f5593, plain, (~ p305(sK86(sK92)) | ~ spl93_237), inference(subsumption_resolution, [], [f4469, f532])).
fof(f4469, plain, (~ p305(sK86(sK92)) | ~ sP2(sK92) | ~ spl93_237), inference(resolution, [], [f1998, f390])).
fof(f390, plain, ! [X0] : (~ p205(sK85(X0)) | ~ p305(sK86(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f210])).
fof(f1998, plain, (p205(sK85(sK92)) | ~ spl93_237), inference(avatar_component_clause, [], [f1996])).
fof(f5574, plain, (~ spl93_8 | ~ spl93_28), inference(avatar_contradiction_clause, [], [f5573])).
fof(f5573, plain, ($false | (~ spl93_8 | ~ spl93_28)), inference(subsumption_resolution, [], [f5572, f550])).
fof(f550, plain, sP27(sK92), inference(resolution, [], [f523, f260])).
fof(f260, plain, ! [X0] : (~ sP40(X0) | sP27(X0)), inference(cnf_transformation, [], [f50])).
fof(f5572, plain, (~ sP27(sK92) | (~ spl93_8 | ~ spl93_28)), inference(subsumption_resolution, [], [f5571, f513])).
fof(f513, plain, (p603(sK92) | ~ spl93_28), inference(avatar_component_clause, [], [f511])).
fof(f511, plain, (spl93_28 <=> p603(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_28])])).
fof(f5571, plain, (~ p603(sK92) | ~ sP27(sK92) | (~ spl93_8 | ~ spl93_28)), inference(resolution, [], [f5564, f324])).
fof(f324, plain, ! [X0] : (~ p203(sK53(X0)) | ~ p603(X0) | ~ sP27(X0)), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ! [X0] : ((~ p203(sK53(X0)) & r1(X0, sK53(X0))) | ~ p603(X0) | ~ sP27(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK53])], [f100, f101])).
fof(f101, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) => (~ p203(sK53(X0)) & r1(X0, sK53(X0)))), introduced(choice_axiom, [])).
fof(f100, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) | ~ p603(X0) | ~ sP27(X0)), inference(rectify, [], [f99])).
fof(f99, plain, ! [X12] : (? [X27] : (~ p203(X27) & r1(X12, X27)) | ~ p603(X12) | ~ sP27(X12)), inference(nnf_transformation, [], [f34])).
fof(f5564, plain, (p203(sK53(sK92)) | (~ spl93_8 | ~ spl93_28)), inference(subsumption_resolution, [], [f5563, f550])).
fof(f5563, plain, (p203(sK53(sK92)) | ~ sP27(sK92) | (~ spl93_8 | ~ spl93_28)), inference(subsumption_resolution, [], [f5511, f513])).
fof(f5511, plain, (p203(sK53(sK92)) | ~ p603(sK92) | ~ sP27(sK92) | ~ spl93_8), inference(resolution, [], [f432, f323])).
fof(f323, plain, ! [X0] : (r1(X0, sK53(X0)) | ~ p603(X0) | ~ sP27(X0)), inference(cnf_transformation, [], [f102])).
fof(f5496, plain, (~ spl93_19 | ~ spl93_24), inference(avatar_split_clause, [], [f589, f494, f473])).
fof(f494, plain, (spl93_24 <=> p502(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_24])])).
fof(f589, plain, (~ p502(sK92) | ~ p402(sK92)), inference(resolution, [], [f271, f523])).
fof(f271, plain, ! [X0] : (~ sP40(X0) | ~ p502(X0) | ~ p402(X0)), inference(cnf_transformation, [], [f50])).
fof(f5489, plain, (~ spl93_20 | ~ spl93_25), inference(avatar_split_clause, [], [f603, f498, f477])).
fof(f498, plain, (spl93_25 <=> p501(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_25])])).
fof(f603, plain, (~ p501(sK92) | ~ p401(sK92)), inference(resolution, [], [f286, f523])).
fof(f286, plain, ! [X0] : (~ sP40(X0) | ~ p501(X0) | ~ p401(X0)), inference(cnf_transformation, [], [f50])).
fof(f5470, plain, (~ spl93_26 | ~ spl93_429), inference(avatar_split_clause, [], [f5469, f5453, f503])).
fof(f5453, plain, (spl93_429 <=> p105(sK64(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_429])])).
fof(f5469, plain, (~ p605(sK92) | ~ spl93_429), inference(subsumption_resolution, [], [f5463, f533])).
fof(f533, plain, sP16(sK92), inference(resolution, [], [f523, f234])).
fof(f234, plain, ! [X0] : (~ sP40(X0) | sP16(X0)), inference(cnf_transformation, [], [f50])).
fof(f5463, plain, (~ p605(sK92) | ~ sP16(sK92) | ~ spl93_429), inference(resolution, [], [f5455, f346])).
fof(f346, plain, ! [X0] : (~ p105(sK64(X0)) | ~ p605(X0) | ~ sP16(X0)), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ! [X0] : ((~ p105(sK64(X0)) & r1(X0, sK64(X0))) | ~ p605(X0) | ~ sP16(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK64])], [f144, f145])).
fof(f145, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) => (~ p105(sK64(X0)) & r1(X0, sK64(X0)))), introduced(choice_axiom, [])).
fof(f144, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) | ~ p605(X0) | ~ sP16(X0)), inference(rectify, [], [f143])).
fof(f143, plain, ! [X12] : (? [X50] : (~ p105(X50) & r1(X12, X50)) | ~ p605(X12) | ~ sP16(X12)), inference(nnf_transformation, [], [f23])).
fof(f5455, plain, (p105(sK64(sK92)) | ~ spl93_429), inference(avatar_component_clause, [], [f5453])).
fof(f5456, plain, (~ spl93_26 | spl93_429 | ~ spl93_1), inference(avatar_split_clause, [], [f5451, f408, f5453, f503])).
fof(f5451, plain, (p105(sK64(sK92)) | ~ p605(sK92) | ~ spl93_1), inference(subsumption_resolution, [], [f5295, f533])).
fof(f5295, plain, (p105(sK64(sK92)) | ~ p605(sK92) | ~ sP16(sK92) | ~ spl93_1), inference(resolution, [], [f409, f345])).
fof(f345, plain, ! [X0] : (r1(X0, sK64(X0)) | ~ p605(X0) | ~ sP16(X0)), inference(cnf_transformation, [], [f146])).
fof(f5384, plain, (~ spl93_10 | ~ spl93_20), inference(avatar_split_clause, [], [f611, f477, f438])).
fof(f438, plain, (spl93_10 <=> p201(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_10])])).
fof(f611, plain, (~ p401(sK92) | ~ p201(sK92)), inference(resolution, [], [f292, f523])).
fof(f292, plain, ! [X0] : (~ sP40(X0) | ~ p401(X0) | ~ p201(X0)), inference(cnf_transformation, [], [f50])).
fof(f5377, plain, (~ spl93_1 | ~ spl93_21), inference(avatar_contradiction_clause, [], [f5376])).
fof(f5376, plain, ($false | (~ spl93_1 | ~ spl93_21)), inference(subsumption_resolution, [], [f5375, f534])).
fof(f534, plain, sP17(sK92), inference(resolution, [], [f523, f235])).
fof(f235, plain, ! [X0] : (~ sP40(X0) | sP17(X0)), inference(cnf_transformation, [], [f50])).
fof(f5375, plain, (~ sP17(sK92) | (~ spl93_1 | ~ spl93_21)), inference(subsumption_resolution, [], [f5374, f484])).
fof(f5374, plain, (~ p505(sK92) | ~ sP17(sK92) | (~ spl93_1 | ~ spl93_21)), inference(resolution, [], [f5339, f344])).
fof(f344, plain, ! [X0] : (~ p105(sK63(X0)) | ~ p505(X0) | ~ sP17(X0)), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ! [X0] : ((~ p105(sK63(X0)) & r1(X0, sK63(X0))) | ~ p505(X0) | ~ sP17(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK63])], [f140, f141])).
fof(f141, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) => (~ p105(sK63(X0)) & r1(X0, sK63(X0)))), introduced(choice_axiom, [])).
fof(f140, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) | ~ p505(X0) | ~ sP17(X0)), inference(rectify, [], [f139])).
fof(f139, plain, ! [X12] : (? [X49] : (~ p105(X49) & r1(X12, X49)) | ~ p505(X12) | ~ sP17(X12)), inference(nnf_transformation, [], [f24])).
fof(f5339, plain, (p105(sK63(sK92)) | (~ spl93_1 | ~ spl93_21)), inference(subsumption_resolution, [], [f5338, f534])).
fof(f5338, plain, (p105(sK63(sK92)) | ~ sP17(sK92) | (~ spl93_1 | ~ spl93_21)), inference(subsumption_resolution, [], [f5294, f484])).
fof(f5294, plain, (p105(sK63(sK92)) | ~ p505(sK92) | ~ sP17(sK92) | ~ spl93_1), inference(resolution, [], [f409, f343])).
fof(f343, plain, ! [X0] : (r1(X0, sK63(X0)) | ~ p505(X0) | ~ sP17(X0)), inference(cnf_transformation, [], [f142])).
fof(f5367, plain, (~ spl93_203 | ~ spl93_1 | spl93_208), inference(avatar_split_clause, [], [f5366, f1812, f408, f1779])).
fof(f1779, plain, (spl93_203 <=> p205(sK80(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_203])])).
fof(f1812, plain, (spl93_208 <=> p105(sK79(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_208])])).
fof(f5366, plain, (~ p205(sK80(sK92)) | (~ spl93_1 | spl93_208)), inference(subsumption_resolution, [], [f5356, f537])).
fof(f537, plain, sP5(sK92), inference(resolution, [], [f523, f238])).
fof(f238, plain, ! [X0] : (~ sP40(X0) | sP5(X0)), inference(cnf_transformation, [], [f50])).
fof(f5356, plain, (~ p205(sK80(sK92)) | ~ sP5(sK92) | (~ spl93_1 | spl93_208)), inference(subsumption_resolution, [], [f5314, f1814])).
fof(f1814, plain, (~ p105(sK79(sK92)) | spl93_208), inference(avatar_component_clause, [], [f1812])).
fof(f5314, plain, (p105(sK79(sK92)) | ~ p205(sK80(sK92)) | ~ sP5(sK92) | ~ spl93_1), inference(resolution, [], [f409, f376])).
fof(f376, plain, ! [X0] : (r1(X0, sK79(X0)) | ~ p205(sK80(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f195])).
fof(f195, plain, ! [X0] : ((~ p105(sK79(X0)) & r1(X0, sK79(X0))) | (~ p205(sK80(X0)) & r1(X0, sK80(X0))) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK79, sK80])], [f192, f194, f193])).
fof(f193, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) => (~ p105(sK79(X0)) & r1(X0, sK79(X0)))), introduced(choice_axiom, [])).
fof(f194, plain, ! [X0] : (? [X2] : (~ p205(X2) & r1(X0, X2)) => (~ p205(sK80(X0)) & r1(X0, sK80(X0)))), introduced(choice_axiom, [])).
fof(f192, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) | ? [X2] : (~ p205(X2) & r1(X0, X2)) | ~ sP5(X0)), inference(rectify, [], [f191])).
fof(f191, plain, ! [X12] : (? [X43] : (~ p105(X43) & r1(X12, X43)) | ? [X44] : (~ p205(X44) & r1(X12, X44)) | ~ sP5(X12)), inference(nnf_transformation, [], [f12])).
fof(f5247, plain, (~ spl93_18 | ~ spl93_23), inference(avatar_split_clause, [], [f580, f490, f469])).
fof(f580, plain, (~ p503(sK92) | ~ p403(sK92)), inference(resolution, [], [f256, f523])).
fof(f256, plain, ! [X0] : (~ sP40(X0) | ~ p503(X0) | ~ p403(X0)), inference(cnf_transformation, [], [f50])).
fof(f5246, plain, (~ spl93_21 | ~ spl93_26), inference(avatar_split_clause, [], [f564, f503, f482])).
fof(f564, plain, (~ p605(sK92) | ~ p505(sK92)), inference(resolution, [], [f224, f523])).
fof(f224, plain, ! [X0] : (~ sP40(X0) | ~ p605(X0) | ~ p505(X0)), inference(cnf_transformation, [], [f50])).
fof(f5221, plain, (~ spl93_17 | ~ spl93_405), inference(avatar_split_clause, [], [f5220, f5200, f465])).
fof(f5200, plain, (spl93_405 <=> p304(sK60(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_405])])).
fof(f5220, plain, (~ p404(sK92) | ~ spl93_405), inference(subsumption_resolution, [], [f5216, f540])).
fof(f540, plain, sP20(sK92), inference(resolution, [], [f523, f244])).
fof(f244, plain, ! [X0] : (~ sP40(X0) | sP20(X0)), inference(cnf_transformation, [], [f50])).
fof(f5216, plain, (~ p404(sK92) | ~ sP20(sK92) | ~ spl93_405), inference(resolution, [], [f5202, f338])).
fof(f338, plain, ! [X0] : (~ p304(sK60(X0)) | ~ p404(X0) | ~ sP20(X0)), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ! [X0] : ((~ p304(sK60(X0)) & r1(X0, sK60(X0))) | ~ p404(X0) | ~ sP20(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK60])], [f128, f129])).
fof(f129, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) => (~ p304(sK60(X0)) & r1(X0, sK60(X0)))), introduced(choice_axiom, [])).
fof(f128, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) | ~ p404(X0) | ~ sP20(X0)), inference(rectify, [], [f127])).
fof(f127, plain, ! [X12] : (? [X40] : (~ p304(X40) & r1(X12, X40)) | ~ p404(X12) | ~ sP20(X12)), inference(nnf_transformation, [], [f27])).
fof(f5202, plain, (p304(sK60(sK92)) | ~ spl93_405), inference(avatar_component_clause, [], [f5200])).
fof(f5203, plain, (~ spl93_17 | spl93_405 | ~ spl93_12), inference(avatar_split_clause, [], [f5198, f446, f5200, f465])).
fof(f446, plain, (spl93_12 <=> ! [X9] : (p304(X9) | ~ r1(sK92, X9))), introduced(avatar_definition, [new_symbols(naming, [spl93_12])])).
fof(f5198, plain, (p304(sK60(sK92)) | ~ p404(sK92) | ~ spl93_12), inference(subsumption_resolution, [], [f4761, f540])).
fof(f4761, plain, (p304(sK60(sK92)) | ~ p404(sK92) | ~ sP20(sK92) | ~ spl93_12), inference(resolution, [], [f447, f337])).
fof(f337, plain, ! [X0] : (r1(X0, sK60(X0)) | ~ p404(X0) | ~ sP20(X0)), inference(cnf_transformation, [], [f130])).
fof(f447, plain, (! [X9] : (~ r1(sK92, X9) | p304(X9)) | ~ spl93_12), inference(avatar_component_clause, [], [f446])).
fof(f5151, plain, (~ spl93_19 | ~ spl93_388), inference(avatar_split_clause, [], [f5150, f5086, f473])).
fof(f5086, plain, (spl93_388 <=> p102(sK43(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_388])])).
fof(f5150, plain, (~ p402(sK92) | ~ spl93_388), inference(subsumption_resolution, [], [f5146, f561])).
fof(f561, plain, sP37(sK92), inference(resolution, [], [f523, f281])).
fof(f281, plain, ! [X0] : (~ sP40(X0) | sP37(X0)), inference(cnf_transformation, [], [f50])).
fof(f5146, plain, (~ p402(sK92) | ~ sP37(sK92) | ~ spl93_388), inference(resolution, [], [f5088, f304])).
fof(f304, plain, ! [X0] : (~ p102(sK43(X0)) | ~ p402(X0) | ~ sP37(X0)), inference(cnf_transformation, [], [f62])).
fof(f62, plain, ! [X0] : ((~ p102(sK43(X0)) & r1(X0, sK43(X0))) | ~ p402(X0) | ~ sP37(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK43])], [f60, f61])).
fof(f61, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) => (~ p102(sK43(X0)) & r1(X0, sK43(X0)))), introduced(choice_axiom, [])).
fof(f60, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) | ~ p402(X0) | ~ sP37(X0)), inference(rectify, [], [f59])).
fof(f59, plain, ! [X12] : (? [X15] : (~ p102(X15) & r1(X12, X15)) | ~ p402(X12) | ~ sP37(X12)), inference(nnf_transformation, [], [f44])).
fof(f5088, plain, (p102(sK43(sK92)) | ~ spl93_388), inference(avatar_component_clause, [], [f5086])).
fof(f5145, plain, (~ spl93_27 | ~ spl93_381), inference(avatar_split_clause, [], [f5144, f5040, f507])).
fof(f5040, plain, (spl93_381 <=> p304(sK62(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_381])])).
fof(f5144, plain, (~ p604(sK92) | ~ spl93_381), inference(subsumption_resolution, [], [f5080, f538])).
fof(f538, plain, sP18(sK92), inference(resolution, [], [f523, f242])).
fof(f242, plain, ! [X0] : (~ sP40(X0) | sP18(X0)), inference(cnf_transformation, [], [f50])).
fof(f5080, plain, (~ p604(sK92) | ~ sP18(sK92) | ~ spl93_381), inference(resolution, [], [f5042, f342])).
fof(f342, plain, ! [X0] : (~ p304(sK62(X0)) | ~ p604(X0) | ~ sP18(X0)), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ! [X0] : ((~ p304(sK62(X0)) & r1(X0, sK62(X0))) | ~ p604(X0) | ~ sP18(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK62])], [f136, f137])).
fof(f137, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) => (~ p304(sK62(X0)) & r1(X0, sK62(X0)))), introduced(choice_axiom, [])).
fof(f136, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) | ~ p604(X0) | ~ sP18(X0)), inference(rectify, [], [f135])).
fof(f135, plain, ! [X12] : (? [X42] : (~ p304(X42) & r1(X12, X42)) | ~ p604(X12) | ~ sP18(X12)), inference(nnf_transformation, [], [f25])).
fof(f5042, plain, (p304(sK62(sK92)) | ~ spl93_381), inference(avatar_component_clause, [], [f5040])).
fof(f5089, plain, (~ spl93_19 | spl93_388 | ~ spl93_4), inference(avatar_split_clause, [], [f5084, f417, f5086, f473])).
fof(f5084, plain, (p102(sK43(sK92)) | ~ p402(sK92) | ~ spl93_4), inference(subsumption_resolution, [], [f4861, f561])).
fof(f4861, plain, (p102(sK43(sK92)) | ~ p402(sK92) | ~ sP37(sK92) | ~ spl93_4), inference(resolution, [], [f418, f303])).
fof(f303, plain, ! [X0] : (r1(X0, sK43(X0)) | ~ p402(X0) | ~ sP37(X0)), inference(cnf_transformation, [], [f62])).
fof(f5043, plain, (~ spl93_27 | spl93_381 | ~ spl93_12), inference(avatar_split_clause, [], [f5038, f446, f5040, f507])).
fof(f5038, plain, (p304(sK62(sK92)) | ~ p604(sK92) | ~ spl93_12), inference(subsumption_resolution, [], [f4763, f538])).
fof(f4763, plain, (p304(sK62(sK92)) | ~ p604(sK92) | ~ sP18(sK92) | ~ spl93_12), inference(resolution, [], [f447, f341])).
fof(f341, plain, ! [X0] : (r1(X0, sK62(X0)) | ~ p604(X0) | ~ sP18(X0)), inference(cnf_transformation, [], [f138])).
fof(f5019, plain, (~ spl93_22 | ~ spl93_27), inference(avatar_split_clause, [], [f569, f507, f486])).
fof(f569, plain, (~ p604(sK92) | ~ p504(sK92)), inference(resolution, [], [f239, f523])).
fof(f239, plain, ! [X0] : (~ sP40(X0) | ~ p604(X0) | ~ p504(X0)), inference(cnf_transformation, [], [f50])).
fof(f4994, plain, (~ spl93_22 | ~ spl93_370), inference(avatar_split_clause, [], [f4993, f4967, f486])).
fof(f4967, plain, (spl93_370 <=> p304(sK61(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_370])])).
fof(f4993, plain, (~ p504(sK92) | ~ spl93_370), inference(subsumption_resolution, [], [f4989, f539])).
fof(f539, plain, sP19(sK92), inference(resolution, [], [f523, f243])).
fof(f243, plain, ! [X0] : (~ sP40(X0) | sP19(X0)), inference(cnf_transformation, [], [f50])).
fof(f4989, plain, (~ p504(sK92) | ~ sP19(sK92) | ~ spl93_370), inference(resolution, [], [f4969, f340])).
fof(f340, plain, ! [X0] : (~ p304(sK61(X0)) | ~ p504(X0) | ~ sP19(X0)), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ! [X0] : ((~ p304(sK61(X0)) & r1(X0, sK61(X0))) | ~ p504(X0) | ~ sP19(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK61])], [f132, f133])).
fof(f133, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) => (~ p304(sK61(X0)) & r1(X0, sK61(X0)))), introduced(choice_axiom, [])).
fof(f132, plain, ! [X0] : (? [X1] : (~ p304(X1) & r1(X0, X1)) | ~ p504(X0) | ~ sP19(X0)), inference(rectify, [], [f131])).
fof(f131, plain, ! [X12] : (? [X41] : (~ p304(X41) & r1(X12, X41)) | ~ p504(X12) | ~ sP19(X12)), inference(nnf_transformation, [], [f26])).
fof(f4969, plain, (p304(sK61(sK92)) | ~ spl93_370), inference(avatar_component_clause, [], [f4967])).
fof(f4970, plain, (~ spl93_22 | spl93_370 | ~ spl93_12), inference(avatar_split_clause, [], [f4965, f446, f4967, f486])).
fof(f4965, plain, (p304(sK61(sK92)) | ~ p504(sK92) | ~ spl93_12), inference(subsumption_resolution, [], [f4762, f539])).
fof(f4762, plain, (p304(sK61(sK92)) | ~ p504(sK92) | ~ sP19(sK92) | ~ spl93_12), inference(resolution, [], [f447, f339])).
fof(f339, plain, ! [X0] : (r1(X0, sK61(X0)) | ~ p504(X0) | ~ sP19(X0)), inference(cnf_transformation, [], [f134])).
fof(f4946, plain, (~ spl93_4 | ~ spl93_24), inference(avatar_contradiction_clause, [], [f4945])).
fof(f4945, plain, ($false | (~ spl93_4 | ~ spl93_24)), inference(subsumption_resolution, [], [f4944, f560])).
fof(f560, plain, sP36(sK92), inference(resolution, [], [f523, f280])).
fof(f280, plain, ! [X0] : (~ sP40(X0) | sP36(X0)), inference(cnf_transformation, [], [f50])).
fof(f4944, plain, (~ sP36(sK92) | (~ spl93_4 | ~ spl93_24)), inference(subsumption_resolution, [], [f4943, f496])).
fof(f496, plain, (p502(sK92) | ~ spl93_24), inference(avatar_component_clause, [], [f494])).
fof(f4943, plain, (~ p502(sK92) | ~ sP36(sK92) | (~ spl93_4 | ~ spl93_24)), inference(resolution, [], [f4920, f306])).
fof(f306, plain, ! [X0] : (~ p102(sK44(X0)) | ~ p502(X0) | ~ sP36(X0)), inference(cnf_transformation, [], [f66])).
fof(f66, plain, ! [X0] : ((~ p102(sK44(X0)) & r1(X0, sK44(X0))) | ~ p502(X0) | ~ sP36(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK44])], [f64, f65])).
fof(f65, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) => (~ p102(sK44(X0)) & r1(X0, sK44(X0)))), introduced(choice_axiom, [])).
fof(f64, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) | ~ p502(X0) | ~ sP36(X0)), inference(rectify, [], [f63])).
fof(f63, plain, ! [X12] : (? [X16] : (~ p102(X16) & r1(X12, X16)) | ~ p502(X12) | ~ sP36(X12)), inference(nnf_transformation, [], [f43])).
fof(f4920, plain, (p102(sK44(sK92)) | (~ spl93_4 | ~ spl93_24)), inference(subsumption_resolution, [], [f4919, f560])).
fof(f4919, plain, (p102(sK44(sK92)) | ~ sP36(sK92) | (~ spl93_4 | ~ spl93_24)), inference(subsumption_resolution, [], [f4862, f496])).
fof(f4862, plain, (p102(sK44(sK92)) | ~ p502(sK92) | ~ sP36(sK92) | ~ spl93_4), inference(resolution, [], [f418, f305])).
fof(f305, plain, ! [X0] : (r1(X0, sK44(X0)) | ~ p502(X0) | ~ sP36(X0)), inference(cnf_transformation, [], [f66])).
fof(f4844, plain, (~ spl93_178 | ~ spl93_181), inference(avatar_contradiction_clause, [], [f4843])).
fof(f4843, plain, ($false | (~ spl93_178 | ~ spl93_181)), inference(subsumption_resolution, [], [f4842, f548])).
fof(f548, plain, sP7(sK92), inference(resolution, [], [f523, f252])).
fof(f252, plain, ! [X0] : (~ sP40(X0) | sP7(X0)), inference(cnf_transformation, [], [f50])).
fof(f4842, plain, (~ sP7(sK92) | (~ spl93_178 | ~ spl93_181)), inference(subsumption_resolution, [], [f4841, f1626])).
fof(f1626, plain, (p304(sK76(sK92)) | ~ spl93_178), inference(avatar_component_clause, [], [f1625])).
fof(f1625, plain, (spl93_178 <=> p304(sK76(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_178])])).
fof(f4841, plain, (~ p304(sK76(sK92)) | ~ sP7(sK92) | ~ spl93_181), inference(resolution, [], [f1645, f370])).
fof(f370, plain, ! [X0] : (~ p104(sK75(X0)) | ~ p304(sK76(X0)) | ~ sP7(X0)), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ! [X0] : ((~ p104(sK75(X0)) & r1(X0, sK75(X0))) | (~ p304(sK76(X0)) & r1(X0, sK76(X0))) | ~ sP7(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK75, sK76])], [f182, f184, f183])).
fof(f183, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) => (~ p104(sK75(X0)) & r1(X0, sK75(X0)))), introduced(choice_axiom, [])).
fof(f184, plain, ! [X0] : (? [X2] : (~ p304(X2) & r1(X0, X2)) => (~ p304(sK76(X0)) & r1(X0, sK76(X0)))), introduced(choice_axiom, [])).
fof(f182, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) | ? [X2] : (~ p304(X2) & r1(X0, X2)) | ~ sP7(X0)), inference(rectify, [], [f181])).
fof(f181, plain, ! [X12] : (? [X30] : (~ p104(X30) & r1(X12, X30)) | ? [X31] : (~ p304(X31) & r1(X12, X31)) | ~ sP7(X12)), inference(nnf_transformation, [], [f14])).
fof(f1645, plain, (p104(sK75(sK92)) | ~ spl93_181), inference(avatar_component_clause, [], [f1643])).
fof(f1643, plain, (spl93_181 <=> p104(sK75(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_181])])).
fof(f4838, plain, (spl93_280 | spl93_178 | ~ spl93_12), inference(avatar_split_clause, [], [f4837, f446, f1625, f2748])).
fof(f2748, plain, (spl93_280 <=> r1(sK92, sK75(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_280])])).
fof(f4837, plain, (p304(sK76(sK92)) | r1(sK92, sK75(sK92)) | ~ spl93_12), inference(subsumption_resolution, [], [f4779, f548])).
fof(f4779, plain, (p304(sK76(sK92)) | r1(sK92, sK75(sK92)) | ~ sP7(sK92) | ~ spl93_12), inference(resolution, [], [f447, f367])).
fof(f367, plain, ! [X0] : (r1(X0, sK76(X0)) | r1(X0, sK75(X0)) | ~ sP7(X0)), inference(cnf_transformation, [], [f185])).
fof(f4834, plain, (~ spl93_181 | spl93_178 | ~ spl93_12), inference(avatar_split_clause, [], [f4833, f446, f1625, f1643])).
fof(f4833, plain, (p304(sK76(sK92)) | ~ p104(sK75(sK92)) | ~ spl93_12), inference(subsumption_resolution, [], [f4780, f548])).
fof(f4780, plain, (p304(sK76(sK92)) | ~ p104(sK75(sK92)) | ~ sP7(sK92) | ~ spl93_12), inference(resolution, [], [f447, f369])).
fof(f369, plain, ! [X0] : (r1(X0, sK76(X0)) | ~ p104(sK75(X0)) | ~ sP7(X0)), inference(cnf_transformation, [], [f185])).
fof(f4832, plain, (~ spl93_178 | spl93_181 | ~ spl93_2), inference(avatar_split_clause, [], [f4831, f411, f1643, f1625])).
fof(f4831, plain, (p104(sK75(sK92)) | ~ p304(sK76(sK92)) | ~ spl93_2), inference(subsumption_resolution, [], [f4712, f548])).
fof(f4712, plain, (p104(sK75(sK92)) | ~ p304(sK76(sK92)) | ~ sP7(sK92) | ~ spl93_2), inference(resolution, [], [f412, f368])).
fof(f368, plain, ! [X0] : (r1(X0, sK75(X0)) | ~ p304(sK76(X0)) | ~ sP7(X0)), inference(cnf_transformation, [], [f185])).
fof(f4824, plain, (~ spl93_2 | spl93_181 | ~ spl93_280), inference(avatar_contradiction_clause, [], [f4823])).
fof(f4823, plain, ($false | (~ spl93_2 | spl93_181 | ~ spl93_280)), inference(subsumption_resolution, [], [f4821, f1644])).
fof(f1644, plain, (~ p104(sK75(sK92)) | spl93_181), inference(avatar_component_clause, [], [f1643])).
fof(f4821, plain, (p104(sK75(sK92)) | (~ spl93_2 | ~ spl93_280)), inference(resolution, [], [f2750, f412])).
fof(f2750, plain, (r1(sK92, sK75(sK92)) | ~ spl93_280), inference(avatar_component_clause, [], [f2748])).
fof(f4673, plain, (~ spl93_4 | ~ spl93_9), inference(avatar_contradiction_clause, [], [f4672])).
fof(f4672, plain, ($false | (~ spl93_4 | ~ spl93_9)), inference(subsumption_resolution, [], [f4671, f563])).
fof(f563, plain, sP39(sK92), inference(resolution, [], [f523, f283])).
fof(f283, plain, ! [X0] : (~ sP40(X0) | sP39(X0)), inference(cnf_transformation, [], [f50])).
fof(f4671, plain, (~ sP39(sK92) | (~ spl93_4 | ~ spl93_9)), inference(subsumption_resolution, [], [f4670, f436])).
fof(f436, plain, (p202(sK92) | ~ spl93_9), inference(avatar_component_clause, [], [f434])).
fof(f434, plain, (spl93_9 <=> p202(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_9])])).
fof(f4670, plain, (~ p202(sK92) | ~ sP39(sK92) | (~ spl93_4 | ~ spl93_9)), inference(resolution, [], [f4641, f300])).
fof(f300, plain, ! [X0] : (~ p102(sK41(X0)) | ~ p202(X0) | ~ sP39(X0)), inference(cnf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : ((~ p102(sK41(X0)) & r1(X0, sK41(X0))) | ~ p202(X0) | ~ sP39(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK41])], [f52, f53])).
fof(f53, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) => (~ p102(sK41(X0)) & r1(X0, sK41(X0)))), introduced(choice_axiom, [])).
fof(f52, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) | ~ p202(X0) | ~ sP39(X0)), inference(rectify, [], [f51])).
fof(f51, plain, ! [X12] : (? [X13] : (~ p102(X13) & r1(X12, X13)) | ~ p202(X12) | ~ sP39(X12)), inference(nnf_transformation, [], [f46])).
fof(f4641, plain, (p102(sK41(sK92)) | (~ spl93_4 | ~ spl93_9)), inference(subsumption_resolution, [], [f4640, f563])).
fof(f4640, plain, (p102(sK41(sK92)) | ~ sP39(sK92) | (~ spl93_4 | ~ spl93_9)), inference(subsumption_resolution, [], [f4580, f436])).
fof(f4580, plain, (p102(sK41(sK92)) | ~ p202(sK92) | ~ sP39(sK92) | ~ spl93_4), inference(resolution, [], [f418, f299])).
fof(f299, plain, ! [X0] : (r1(X0, sK41(X0)) | ~ p202(X0) | ~ sP39(X0)), inference(cnf_transformation, [], [f54])).
fof(f4577, plain, (~ spl93_27 | ~ spl93_145), inference(avatar_contradiction_clause, [], [f4576])).
fof(f4576, plain, ($false | (~ spl93_27 | ~ spl93_145)), inference(subsumption_resolution, [], [f4575, f545])).
fof(f545, plain, sP24(sK92), inference(resolution, [], [f523, f249])).
fof(f249, plain, ! [X0] : (~ sP40(X0) | sP24(X0)), inference(cnf_transformation, [], [f50])).
fof(f4575, plain, (~ sP24(sK92) | (~ spl93_27 | ~ spl93_145)), inference(subsumption_resolution, [], [f4574, f509])).
fof(f4574, plain, (~ p604(sK92) | ~ sP24(sK92) | ~ spl93_145), inference(resolution, [], [f1427, f330])).
fof(f330, plain, ! [X0] : (~ p104(sK56(X0)) | ~ p604(X0) | ~ sP24(X0)), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ! [X0] : ((~ p104(sK56(X0)) & r1(X0, sK56(X0))) | ~ p604(X0) | ~ sP24(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK56])], [f112, f113])).
fof(f113, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) => (~ p104(sK56(X0)) & r1(X0, sK56(X0)))), introduced(choice_axiom, [])).
fof(f112, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) | ~ p604(X0) | ~ sP24(X0)), inference(rectify, [], [f111])).
fof(f111, plain, ! [X12] : (? [X34] : (~ p104(X34) & r1(X12, X34)) | ~ p604(X12) | ~ sP24(X12)), inference(nnf_transformation, [], [f31])).
fof(f1427, plain, (p104(sK56(sK92)) | ~ spl93_145), inference(avatar_component_clause, [], [f1425])).
fof(f1425, plain, (spl93_145 <=> p104(sK56(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_145])])).
fof(f4495, plain, (~ spl93_9 | ~ spl93_24), inference(avatar_split_clause, [], [f596, f494, f434])).
fof(f596, plain, (~ p502(sK92) | ~ p202(sK92)), inference(resolution, [], [f276, f523])).
fof(f276, plain, ! [X0] : (~ sP40(X0) | ~ p502(X0) | ~ p202(X0)), inference(cnf_transformation, [], [f50])).
fof(f4494, plain, (~ spl93_153 | spl93_158 | ~ spl93_3), inference(avatar_split_clause, [], [f4493, f414, f1504, f1471])).
fof(f1471, plain, (spl93_153 <=> p203(sK72(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_153])])).
fof(f1504, plain, (spl93_158 <=> p103(sK71(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_158])])).
fof(f4493, plain, (p103(sK71(sK92)) | ~ p203(sK72(sK92)) | ~ spl93_3), inference(subsumption_resolution, [], [f4092, f558])).
fof(f558, plain, sP9(sK92), inference(resolution, [], [f523, f268])).
fof(f268, plain, ! [X0] : (~ sP40(X0) | sP9(X0)), inference(cnf_transformation, [], [f50])).
fof(f4092, plain, (p103(sK71(sK92)) | ~ p203(sK72(sK92)) | ~ sP9(sK92) | ~ spl93_3), inference(resolution, [], [f415, f360])).
fof(f360, plain, ! [X0] : (r1(X0, sK71(X0)) | ~ p203(sK72(X0)) | ~ sP9(X0)), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ! [X0] : ((~ p103(sK71(X0)) & r1(X0, sK71(X0))) | (~ p203(sK72(X0)) & r1(X0, sK72(X0))) | ~ sP9(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK71, sK72])], [f172, f174, f173])).
fof(f173, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) => (~ p103(sK71(X0)) & r1(X0, sK71(X0)))), introduced(choice_axiom, [])).
fof(f174, plain, ! [X0] : (? [X2] : (~ p203(X2) & r1(X0, X2)) => (~ p203(sK72(X0)) & r1(X0, sK72(X0)))), introduced(choice_axiom, [])).
fof(f172, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) | ? [X2] : (~ p203(X2) & r1(X0, X2)) | ~ sP9(X0)), inference(rectify, [], [f171])).
fof(f171, plain, ! [X12] : (? [X18] : (~ p103(X18) & r1(X12, X18)) | ? [X19] : (~ p203(X19) & r1(X12, X19)) | ~ sP9(X12)), inference(nnf_transformation, [], [f16])).
fof(f4478, plain, (spl93_248 | ~ spl93_6 | ~ spl93_291), inference(avatar_split_clause, [], [f4393, f2816, f425, f2068])).
fof(f2068, plain, (spl93_248 <=> p205(sK87(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_248])])).
fof(f425, plain, (spl93_6 <=> ! [X8] : (p205(X8) | ~ r1(sK92, X8))), introduced(avatar_definition, [new_symbols(naming, [spl93_6])])).
fof(f2816, plain, (spl93_291 <=> r1(sK92, sK87(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_291])])).
fof(f4393, plain, (p205(sK87(sK92)) | (~ spl93_6 | ~ spl93_291)), inference(resolution, [], [f426, f2818])).
fof(f2818, plain, (r1(sK92, sK87(sK92)) | ~ spl93_291), inference(avatar_component_clause, [], [f2816])).
fof(f426, plain, (! [X8] : (~ r1(sK92, X8) | p205(X8)) | ~ spl93_6), inference(avatar_component_clause, [], [f425])).
fof(f4477, plain, (~ spl93_246 | spl93_248 | ~ spl93_6), inference(avatar_split_clause, [], [f4476, f425, f2068, f2056])).
fof(f2056, plain, (spl93_246 <=> p405(sK88(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_246])])).
fof(f4476, plain, (p205(sK87(sK92)) | ~ p405(sK88(sK92)) | ~ spl93_6), inference(subsumption_resolution, [], [f4448, f531])).
fof(f531, plain, sP1(sK92), inference(resolution, [], [f523, f232])).
fof(f232, plain, ! [X0] : (~ sP40(X0) | sP1(X0)), inference(cnf_transformation, [], [f50])).
fof(f4448, plain, (p205(sK87(sK92)) | ~ p405(sK88(sK92)) | ~ sP1(sK92) | ~ spl93_6), inference(resolution, [], [f426, f392])).
fof(f392, plain, ! [X0] : (r1(X0, sK87(X0)) | ~ p405(sK88(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f215])).
fof(f215, plain, ! [X0] : ((~ p205(sK87(X0)) & r1(X0, sK87(X0))) | (~ p405(sK88(X0)) & r1(X0, sK88(X0))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK87, sK88])], [f212, f214, f213])).
fof(f213, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) => (~ p205(sK87(X0)) & r1(X0, sK87(X0)))), introduced(choice_axiom, [])).
fof(f214, plain, ! [X0] : (? [X2] : (~ p405(X2) & r1(X0, X2)) => (~ p405(sK88(X0)) & r1(X0, sK88(X0)))), introduced(choice_axiom, [])).
fof(f212, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) | ? [X2] : (~ p405(X2) & r1(X0, X2)) | ~ sP1(X0)), inference(rectify, [], [f211])).
fof(f211, plain, ! [X12] : (? [X53] : (~ p205(X53) & r1(X12, X53)) | ? [X54] : (~ p405(X54) & r1(X12, X54)) | ~ sP1(X12)), inference(nnf_transformation, [], [f8])).
fof(f4475, plain, (~ spl93_248 | spl93_246 | ~ spl93_16), inference(avatar_split_clause, [], [f4474, f462, f2056, f2068])).
fof(f462, plain, (spl93_16 <=> ! [X11] : (p405(X11) | ~ r1(sK92, X11))), introduced(avatar_definition, [new_symbols(naming, [spl93_16])])).
fof(f4474, plain, (p405(sK88(sK92)) | ~ p205(sK87(sK92)) | ~ spl93_16), inference(subsumption_resolution, [], [f3238, f531])).
fof(f3238, plain, (p405(sK88(sK92)) | ~ p205(sK87(sK92)) | ~ sP1(sK92) | ~ spl93_16), inference(resolution, [], [f463, f393])).
fof(f393, plain, ! [X0] : (r1(X0, sK88(X0)) | ~ p205(sK87(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f215])).
fof(f463, plain, (! [X11] : (~ r1(sK92, X11) | p405(X11)) | ~ spl93_16), inference(avatar_component_clause, [], [f462])).
fof(f4473, plain, (~ spl93_246 | ~ spl93_248), inference(avatar_contradiction_clause, [], [f4472])).
fof(f4472, plain, ($false | (~ spl93_246 | ~ spl93_248)), inference(subsumption_resolution, [], [f4471, f531])).
fof(f4471, plain, (~ sP1(sK92) | (~ spl93_246 | ~ spl93_248)), inference(subsumption_resolution, [], [f4470, f2057])).
fof(f2057, plain, (p405(sK88(sK92)) | ~ spl93_246), inference(avatar_component_clause, [], [f2056])).
fof(f4470, plain, (~ p405(sK88(sK92)) | ~ sP1(sK92) | ~ spl93_248), inference(resolution, [], [f2070, f394])).
fof(f394, plain, ! [X0] : (~ p205(sK87(X0)) | ~ p405(sK88(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f215])).
fof(f2070, plain, (p205(sK87(sK92)) | ~ spl93_248), inference(avatar_component_clause, [], [f2068])).
fof(f4460, plain, (spl93_237 | ~ spl93_6 | ~ spl93_289), inference(avatar_split_clause, [], [f4392, f2804, f425, f1996])).
fof(f4392, plain, (p205(sK85(sK92)) | (~ spl93_6 | ~ spl93_289)), inference(resolution, [], [f426, f2806])).
fof(f2806, plain, (r1(sK92, sK85(sK92)) | ~ spl93_289), inference(avatar_component_clause, [], [f2804])).
fof(f4371, plain, (~ spl93_203 | ~ spl93_208), inference(avatar_split_clause, [], [f4370, f1812, f1779])).
fof(f4370, plain, (~ p205(sK80(sK92)) | ~ spl93_208), inference(subsumption_resolution, [], [f3914, f537])).
fof(f3914, plain, (~ p205(sK80(sK92)) | ~ sP5(sK92) | ~ spl93_208), inference(resolution, [], [f1813, f378])).
fof(f378, plain, ! [X0] : (~ p105(sK79(X0)) | ~ p205(sK80(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f195])).
fof(f1813, plain, (p105(sK79(sK92)) | ~ spl93_208), inference(avatar_component_clause, [], [f1812])).
fof(f4363, plain, (~ spl93_153 | ~ spl93_158), inference(avatar_split_clause, [], [f4362, f1504, f1471])).
fof(f4362, plain, (~ p203(sK72(sK92)) | ~ spl93_158), inference(subsumption_resolution, [], [f4138, f558])).
fof(f4138, plain, (~ p203(sK72(sK92)) | ~ sP9(sK92) | ~ spl93_158), inference(resolution, [], [f1505, f362])).
fof(f362, plain, ! [X0] : (~ p103(sK71(X0)) | ~ p203(sK72(X0)) | ~ sP9(X0)), inference(cnf_transformation, [], [f175])).
fof(f1505, plain, (p103(sK71(sK92)) | ~ spl93_158), inference(avatar_component_clause, [], [f1504])).
fof(f4361, plain, (spl93_153 | ~ spl93_8 | ~ spl93_158), inference(avatar_split_clause, [], [f4360, f1504, f431, f1471])).
fof(f4360, plain, (p203(sK72(sK92)) | (~ spl93_8 | ~ spl93_158)), inference(subsumption_resolution, [], [f4359, f558])).
fof(f4359, plain, (p203(sK72(sK92)) | ~ sP9(sK92) | (~ spl93_8 | ~ spl93_158)), inference(subsumption_resolution, [], [f4315, f1505])).
fof(f4315, plain, (p203(sK72(sK92)) | ~ p103(sK71(sK92)) | ~ sP9(sK92) | ~ spl93_8), inference(resolution, [], [f432, f361])).
fof(f361, plain, ! [X0] : (r1(X0, sK72(X0)) | ~ p103(sK71(X0)) | ~ sP9(X0)), inference(cnf_transformation, [], [f175])).
fof(f4254, plain, (~ spl93_10 | ~ spl93_15), inference(avatar_split_clause, [], [f4253, f457, f438])).
fof(f4253, plain, (~ p201(sK92) | ~ spl93_15), inference(subsumption_resolution, [], [f612, f459])).
fof(f459, plain, (p301(sK92) | ~ spl93_15), inference(avatar_component_clause, [], [f457])).
fof(f612, plain, (~ p301(sK92) | ~ p201(sK92)), inference(resolution, [], [f293, f523])).
fof(f293, plain, ! [X0] : (~ sP40(X0) | ~ p301(X0) | ~ p201(X0)), inference(cnf_transformation, [], [f50])).
fof(f4252, plain, (~ spl93_10 | ~ spl93_25), inference(avatar_split_clause, [], [f610, f498, f438])).
fof(f610, plain, (~ p501(sK92) | ~ p201(sK92)), inference(resolution, [], [f291, f523])).
fof(f291, plain, ! [X0] : (~ sP40(X0) | ~ p501(X0) | ~ p201(X0)), inference(cnf_transformation, [], [f50])).
fof(f4251, plain, (~ spl93_23 | ~ spl93_351), inference(avatar_split_clause, [], [f4250, f4241, f490])).
fof(f4241, plain, (spl93_351 <=> p103(sK48(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_351])])).
fof(f4250, plain, (~ p503(sK92) | ~ spl93_351), inference(subsumption_resolution, [], [f4245, f555])).
fof(f555, plain, sP32(sK92), inference(resolution, [], [f523, f265])).
fof(f265, plain, ! [X0] : (~ sP40(X0) | sP32(X0)), inference(cnf_transformation, [], [f50])).
fof(f4245, plain, (~ p503(sK92) | ~ sP32(sK92) | ~ spl93_351), inference(resolution, [], [f4243, f314])).
fof(f314, plain, ! [X0] : (~ p103(sK48(X0)) | ~ p503(X0) | ~ sP32(X0)), inference(cnf_transformation, [], [f82])).
fof(f82, plain, ! [X0] : ((~ p103(sK48(X0)) & r1(X0, sK48(X0))) | ~ p503(X0) | ~ sP32(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK48])], [f80, f81])).
fof(f81, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) => (~ p103(sK48(X0)) & r1(X0, sK48(X0)))), introduced(choice_axiom, [])).
fof(f80, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) | ~ p503(X0) | ~ sP32(X0)), inference(rectify, [], [f79])).
fof(f79, plain, ! [X12] : (? [X22] : (~ p103(X22) & r1(X12, X22)) | ~ p503(X12) | ~ sP32(X12)), inference(nnf_transformation, [], [f39])).
fof(f4243, plain, (p103(sK48(sK92)) | ~ spl93_351), inference(avatar_component_clause, [], [f4241])).
fof(f4249, plain, (~ spl93_15 | ~ spl93_25), inference(avatar_split_clause, [], [f606, f498, f457])).
fof(f606, plain, (~ p501(sK92) | ~ p301(sK92)), inference(resolution, [], [f288, f523])).
fof(f288, plain, ! [X0] : (~ sP40(X0) | ~ p501(X0) | ~ p301(X0)), inference(cnf_transformation, [], [f50])).
fof(f4244, plain, (~ spl93_23 | spl93_351 | ~ spl93_3), inference(avatar_split_clause, [], [f4239, f414, f4241, f490])).
fof(f4239, plain, (p103(sK48(sK92)) | ~ p503(sK92) | ~ spl93_3), inference(subsumption_resolution, [], [f4069, f555])).
fof(f4069, plain, (p103(sK48(sK92)) | ~ p503(sK92) | ~ sP32(sK92) | ~ spl93_3), inference(resolution, [], [f415, f313])).
fof(f313, plain, ! [X0] : (r1(X0, sK48(X0)) | ~ p503(X0) | ~ sP32(X0)), inference(cnf_transformation, [], [f82])).
fof(f4232, plain, (~ spl93_28 | ~ spl93_343), inference(avatar_split_clause, [], [f4231, f4181, f511])).
fof(f4181, plain, (spl93_343 <=> p103(sK49(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_343])])).
fof(f4231, plain, (~ p603(sK92) | ~ spl93_343), inference(subsumption_resolution, [], [f4211, f554])).
fof(f554, plain, sP31(sK92), inference(resolution, [], [f523, f264])).
fof(f264, plain, ! [X0] : (~ sP40(X0) | sP31(X0)), inference(cnf_transformation, [], [f50])).
fof(f4211, plain, (~ p603(sK92) | ~ sP31(sK92) | ~ spl93_343), inference(resolution, [], [f4183, f316])).
fof(f316, plain, ! [X0] : (~ p103(sK49(X0)) | ~ p603(X0) | ~ sP31(X0)), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ! [X0] : ((~ p103(sK49(X0)) & r1(X0, sK49(X0))) | ~ p603(X0) | ~ sP31(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK49])], [f84, f85])).
fof(f85, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) => (~ p103(sK49(X0)) & r1(X0, sK49(X0)))), introduced(choice_axiom, [])).
fof(f84, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) | ~ p603(X0) | ~ sP31(X0)), inference(rectify, [], [f83])).
fof(f83, plain, ! [X12] : (? [X23] : (~ p103(X23) & r1(X12, X23)) | ~ p603(X12) | ~ sP31(X12)), inference(nnf_transformation, [], [f38])).
fof(f4183, plain, (p103(sK49(sK92)) | ~ spl93_343), inference(avatar_component_clause, [], [f4181])).
fof(f4230, plain, (~ spl93_24 | ~ spl93_14), inference(avatar_split_clause, [], [f4229, f453, f494])).
fof(f4229, plain, (~ p502(sK92) | ~ spl93_14), inference(subsumption_resolution, [], [f592, f455])).
fof(f592, plain, (~ p502(sK92) | ~ p302(sK92)), inference(resolution, [], [f273, f523])).
fof(f273, plain, ! [X0] : (~ sP40(X0) | ~ p502(X0) | ~ p302(X0)), inference(cnf_transformation, [], [f50])).
fof(f4184, plain, (~ spl93_28 | spl93_343 | ~ spl93_3), inference(avatar_split_clause, [], [f4179, f414, f4181, f511])).
fof(f4179, plain, (p103(sK49(sK92)) | ~ p603(sK92) | ~ spl93_3), inference(subsumption_resolution, [], [f4070, f554])).
fof(f4070, plain, (p103(sK49(sK92)) | ~ p603(sK92) | ~ sP31(sK92) | ~ spl93_3), inference(resolution, [], [f415, f315])).
fof(f315, plain, ! [X0] : (r1(X0, sK49(X0)) | ~ p603(X0) | ~ sP31(X0)), inference(cnf_transformation, [], [f86])).
fof(f4142, plain, (~ spl93_3 | ~ spl93_13), inference(avatar_contradiction_clause, [], [f4141])).
fof(f4141, plain, ($false | (~ spl93_3 | ~ spl93_13)), inference(subsumption_resolution, [], [f4140, f557])).
fof(f557, plain, sP34(sK92), inference(resolution, [], [f523, f267])).
fof(f267, plain, ! [X0] : (~ sP40(X0) | sP34(X0)), inference(cnf_transformation, [], [f50])).
fof(f4140, plain, (~ sP34(sK92) | (~ spl93_3 | ~ spl93_13)), inference(subsumption_resolution, [], [f4139, f451])).
fof(f451, plain, (p303(sK92) | ~ spl93_13), inference(avatar_component_clause, [], [f449])).
fof(f449, plain, (spl93_13 <=> p303(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_13])])).
fof(f4139, plain, (~ p303(sK92) | ~ sP34(sK92) | (~ spl93_3 | ~ spl93_13)), inference(resolution, [], [f4125, f310])).
fof(f310, plain, ! [X0] : (~ p103(sK46(X0)) | ~ p303(X0) | ~ sP34(X0)), inference(cnf_transformation, [], [f74])).
fof(f74, plain, ! [X0] : ((~ p103(sK46(X0)) & r1(X0, sK46(X0))) | ~ p303(X0) | ~ sP34(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK46])], [f72, f73])).
fof(f73, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) => (~ p103(sK46(X0)) & r1(X0, sK46(X0)))), introduced(choice_axiom, [])).
fof(f72, plain, ! [X0] : (? [X1] : (~ p103(X1) & r1(X0, X1)) | ~ p303(X0) | ~ sP34(X0)), inference(rectify, [], [f71])).
fof(f71, plain, ! [X12] : (? [X20] : (~ p103(X20) & r1(X12, X20)) | ~ p303(X12) | ~ sP34(X12)), inference(nnf_transformation, [], [f41])).
fof(f4125, plain, (p103(sK46(sK92)) | (~ spl93_3 | ~ spl93_13)), inference(subsumption_resolution, [], [f4124, f557])).
fof(f4124, plain, (p103(sK46(sK92)) | ~ sP34(sK92) | (~ spl93_3 | ~ spl93_13)), inference(subsumption_resolution, [], [f4067, f451])).
fof(f4067, plain, (p103(sK46(sK92)) | ~ p303(sK92) | ~ sP34(sK92) | ~ spl93_3), inference(resolution, [], [f415, f309])).
fof(f309, plain, ! [X0] : (r1(X0, sK46(X0)) | ~ p303(X0) | ~ sP34(X0)), inference(cnf_transformation, [], [f74])).
fof(f4057, plain, (~ spl93_26 | ~ spl93_124), inference(avatar_split_clause, [], [f4056, f1299, f503])).
fof(f1299, plain, (spl93_124 <=> p405(sK70(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_124])])).
fof(f4056, plain, (~ p605(sK92) | ~ spl93_124), inference(subsumption_resolution, [], [f4052, f524])).
fof(f524, plain, sP10(sK92), inference(resolution, [], [f523, f225])).
fof(f225, plain, ! [X0] : (~ sP40(X0) | sP10(X0)), inference(cnf_transformation, [], [f50])).
fof(f4052, plain, (~ p605(sK92) | ~ sP10(sK92) | ~ spl93_124), inference(resolution, [], [f1301, f358])).
fof(f358, plain, ! [X0] : (~ p405(sK70(X0)) | ~ p605(X0) | ~ sP10(X0)), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ! [X0] : ((~ p405(sK70(X0)) & r1(X0, sK70(X0))) | ~ p605(X0) | ~ sP10(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK70])], [f168, f169])).
fof(f169, plain, ! [X0] : (? [X1] : (~ p405(X1) & r1(X0, X1)) => (~ p405(sK70(X0)) & r1(X0, sK70(X0)))), introduced(choice_axiom, [])).
fof(f168, plain, ! [X0] : (? [X1] : (~ p405(X1) & r1(X0, X1)) | ~ p605(X0) | ~ sP10(X0)), inference(rectify, [], [f167])).
fof(f167, plain, ! [X12] : (? [X62] : (~ p405(X62) & r1(X12, X62)) | ~ p605(X12) | ~ sP10(X12)), inference(nnf_transformation, [], [f17])).
fof(f1301, plain, (p405(sK70(sK92)) | ~ spl93_124), inference(avatar_component_clause, [], [f1299])).
fof(f4021, plain, (~ spl93_26 | spl93_124 | ~ spl93_16), inference(avatar_split_clause, [], [f4020, f462, f1299, f503])).
fof(f4020, plain, (p405(sK70(sK92)) | ~ p605(sK92) | ~ spl93_16), inference(subsumption_resolution, [], [f3211, f524])).
fof(f3211, plain, (p405(sK70(sK92)) | ~ p605(sK92) | ~ sP10(sK92) | ~ spl93_16), inference(resolution, [], [f463, f357])).
fof(f357, plain, ! [X0] : (r1(X0, sK70(X0)) | ~ p605(X0) | ~ sP10(X0)), inference(cnf_transformation, [], [f170])).
fof(f4015, plain, (spl93_158 | ~ spl93_3 | ~ spl93_264), inference(avatar_split_clause, [], [f3925, f2183, f414, f1504])).
fof(f2183, plain, (spl93_264 <=> r1(sK92, sK71(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_264])])).
fof(f3925, plain, (p103(sK71(sK92)) | (~ spl93_3 | ~ spl93_264)), inference(resolution, [], [f415, f2185])).
fof(f2185, plain, (r1(sK92, sK71(sK92)) | ~ spl93_264), inference(avatar_component_clause, [], [f2183])).
fof(f3924, plain, (spl93_287 | spl93_227 | ~ spl93_16), inference(avatar_split_clause, [], [f3923, f462, f1931, f2792])).
fof(f2792, plain, (spl93_287 <=> r1(sK92, sK83(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_287])])).
fof(f1931, plain, (spl93_227 <=> p405(sK84(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_227])])).
fof(f3923, plain, (p405(sK84(sK92)) | r1(sK92, sK83(sK92)) | ~ spl93_16), inference(subsumption_resolution, [], [f3231, f535])).
fof(f535, plain, sP3(sK92), inference(resolution, [], [f523, f236])).
fof(f236, plain, ! [X0] : (~ sP40(X0) | sP3(X0)), inference(cnf_transformation, [], [f50])).
fof(f3231, plain, (p405(sK84(sK92)) | r1(sK92, sK83(sK92)) | ~ sP3(sK92) | ~ spl93_16), inference(resolution, [], [f463, f383])).
fof(f383, plain, ! [X0] : (r1(X0, sK84(X0)) | r1(X0, sK83(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ! [X0] : ((~ p105(sK83(X0)) & r1(X0, sK83(X0))) | (~ p405(sK84(X0)) & r1(X0, sK84(X0))) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK83, sK84])], [f202, f204, f203])).
fof(f203, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) => (~ p105(sK83(X0)) & r1(X0, sK83(X0)))), introduced(choice_axiom, [])).
fof(f204, plain, ! [X0] : (? [X2] : (~ p405(X2) & r1(X0, X2)) => (~ p405(sK84(X0)) & r1(X0, sK84(X0)))), introduced(choice_axiom, [])).
fof(f202, plain, ! [X0] : (? [X1] : (~ p105(X1) & r1(X0, X1)) | ? [X2] : (~ p405(X2) & r1(X0, X2)) | ~ sP3(X0)), inference(rectify, [], [f201])).
fof(f201, plain, ! [X12] : (? [X47] : (~ p105(X47) & r1(X12, X47)) | ? [X48] : (~ p405(X48) & r1(X12, X48)) | ~ sP3(X12)), inference(nnf_transformation, [], [f10])).
fof(f3922, plain, (spl93_266 | ~ spl93_1 | ~ spl93_227), inference(avatar_split_clause, [], [f3921, f1931, f408, f2258])).
fof(f2258, plain, (spl93_266 <=> p105(sK83(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_266])])).
fof(f3921, plain, (p105(sK83(sK92)) | (~ spl93_1 | ~ spl93_227)), inference(subsumption_resolution, [], [f3920, f535])).
fof(f3920, plain, (p105(sK83(sK92)) | ~ sP3(sK92) | (~ spl93_1 | ~ spl93_227)), inference(subsumption_resolution, [], [f3833, f1932])).
fof(f1932, plain, (p405(sK84(sK92)) | ~ spl93_227), inference(avatar_component_clause, [], [f1931])).
fof(f3833, plain, (p105(sK83(sK92)) | ~ p405(sK84(sK92)) | ~ sP3(sK92) | ~ spl93_1), inference(resolution, [], [f409, f384])).
fof(f384, plain, ! [X0] : (r1(X0, sK83(X0)) | ~ p405(sK84(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f205])).
fof(f3919, plain, (~ spl93_227 | ~ spl93_266), inference(avatar_contradiction_clause, [], [f3918])).
fof(f3918, plain, ($false | (~ spl93_227 | ~ spl93_266)), inference(subsumption_resolution, [], [f3917, f535])).
fof(f3917, plain, (~ sP3(sK92) | (~ spl93_227 | ~ spl93_266)), inference(subsumption_resolution, [], [f3916, f1932])).
fof(f3916, plain, (~ p405(sK84(sK92)) | ~ sP3(sK92) | ~ spl93_266), inference(resolution, [], [f2259, f386])).
fof(f386, plain, ! [X0] : (~ p105(sK83(X0)) | ~ p405(sK84(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f205])).
fof(f2259, plain, (p105(sK83(sK92)) | ~ spl93_266), inference(avatar_component_clause, [], [f2258])).
fof(f3907, plain, (spl93_227 | ~ spl93_16 | ~ spl93_266), inference(avatar_split_clause, [], [f3906, f2258, f462, f1931])).
fof(f3906, plain, (p405(sK84(sK92)) | (~ spl93_16 | ~ spl93_266)), inference(subsumption_resolution, [], [f3905, f535])).
fof(f3905, plain, (p405(sK84(sK92)) | ~ sP3(sK92) | (~ spl93_16 | ~ spl93_266)), inference(subsumption_resolution, [], [f3232, f2259])).
fof(f3232, plain, (p405(sK84(sK92)) | ~ p105(sK83(sK92)) | ~ sP3(sK92) | ~ spl93_16), inference(resolution, [], [f463, f385])).
fof(f385, plain, ! [X0] : (r1(X0, sK84(X0)) | ~ p105(sK83(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f205])).
fof(f3850, plain, (~ spl93_1 | spl93_266 | ~ spl93_287), inference(avatar_contradiction_clause, [], [f3849])).
fof(f3849, plain, ($false | (~ spl93_1 | spl93_266 | ~ spl93_287)), inference(subsumption_resolution, [], [f3782, f2260])).
fof(f2260, plain, (~ p105(sK83(sK92)) | spl93_266), inference(avatar_component_clause, [], [f2258])).
fof(f3782, plain, (p105(sK83(sK92)) | (~ spl93_1 | ~ spl93_287)), inference(resolution, [], [f409, f2794])).
fof(f2794, plain, (r1(sK92, sK83(sK92)) | ~ spl93_287), inference(avatar_component_clause, [], [f2792])).
fof(f3848, plain, (~ spl93_1 | spl93_220 | ~ spl93_286), inference(avatar_contradiction_clause, [], [f3847])).
fof(f3847, plain, ($false | (~ spl93_1 | spl93_220 | ~ spl93_286)), inference(subsumption_resolution, [], [f3781, f1890])).
fof(f1890, plain, (~ p105(sK81(sK92)) | spl93_220), inference(avatar_component_clause, [], [f1888])).
fof(f3781, plain, (p105(sK81(sK92)) | (~ spl93_1 | ~ spl93_286)), inference(resolution, [], [f409, f2786])).
fof(f2786, plain, (r1(sK92, sK81(sK92)) | ~ spl93_286), inference(avatar_component_clause, [], [f2784])).
fof(f3846, plain, (~ spl93_1 | spl93_208 | ~ spl93_284), inference(avatar_contradiction_clause, [], [f3845])).
fof(f3845, plain, ($false | (~ spl93_1 | spl93_208 | ~ spl93_284)), inference(subsumption_resolution, [], [f3780, f1814])).
fof(f3780, plain, (p105(sK79(sK92)) | (~ spl93_1 | ~ spl93_284)), inference(resolution, [], [f409, f2774])).
fof(f2774, plain, (r1(sK92, sK79(sK92)) | ~ spl93_284), inference(avatar_component_clause, [], [f2772])).
fof(f2772, plain, (spl93_284 <=> r1(sK92, sK79(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_284])])).
fof(f3777, plain, (~ spl93_4 | ~ spl93_29), inference(avatar_contradiction_clause, [], [f3776])).
fof(f3776, plain, ($false | (~ spl93_4 | ~ spl93_29)), inference(subsumption_resolution, [], [f3775, f559])).
fof(f559, plain, sP35(sK92), inference(resolution, [], [f523, f279])).
fof(f279, plain, ! [X0] : (~ sP40(X0) | sP35(X0)), inference(cnf_transformation, [], [f50])).
fof(f3775, plain, (~ sP35(sK92) | (~ spl93_4 | ~ spl93_29)), inference(subsumption_resolution, [], [f3774, f517])).
fof(f517, plain, (p602(sK92) | ~ spl93_29), inference(avatar_component_clause, [], [f515])).
fof(f515, plain, (spl93_29 <=> p602(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_29])])).
fof(f3774, plain, (~ p602(sK92) | ~ sP35(sK92) | (~ spl93_4 | ~ spl93_29)), inference(resolution, [], [f3763, f308])).
fof(f308, plain, ! [X0] : (~ p102(sK45(X0)) | ~ p602(X0) | ~ sP35(X0)), inference(cnf_transformation, [], [f70])).
fof(f70, plain, ! [X0] : ((~ p102(sK45(X0)) & r1(X0, sK45(X0))) | ~ p602(X0) | ~ sP35(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK45])], [f68, f69])).
fof(f69, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) => (~ p102(sK45(X0)) & r1(X0, sK45(X0)))), introduced(choice_axiom, [])).
fof(f68, plain, ! [X0] : (? [X1] : (~ p102(X1) & r1(X0, X1)) | ~ p602(X0) | ~ sP35(X0)), inference(rectify, [], [f67])).
fof(f67, plain, ! [X12] : (? [X17] : (~ p102(X17) & r1(X12, X17)) | ~ p602(X12) | ~ sP35(X12)), inference(nnf_transformation, [], [f42])).
fof(f3763, plain, (p102(sK45(sK92)) | (~ spl93_4 | ~ spl93_29)), inference(subsumption_resolution, [], [f3762, f559])).
fof(f3762, plain, (p102(sK45(sK92)) | ~ sP35(sK92) | (~ spl93_4 | ~ spl93_29)), inference(subsumption_resolution, [], [f3706, f517])).
fof(f3706, plain, (p102(sK45(sK92)) | ~ p602(sK92) | ~ sP35(sK92) | ~ spl93_4), inference(resolution, [], [f418, f307])).
fof(f307, plain, ! [X0] : (r1(X0, sK45(X0)) | ~ p602(X0) | ~ sP35(X0)), inference(cnf_transformation, [], [f70])).
fof(f3694, plain, (~ spl93_5 | ~ spl93_25), inference(avatar_split_clause, [], [f615, f498, f420])).
fof(f615, plain, (~ p501(sK92) | ~ p101(sK92)), inference(resolution, [], [f295, f523])).
fof(f295, plain, ! [X0] : (~ sP40(X0) | ~ p501(X0) | ~ p101(X0)), inference(cnf_transformation, [], [f50])).
fof(f3693, plain, (~ spl93_190 | ~ spl93_7 | spl93_195), inference(avatar_split_clause, [], [f3692, f1733, f428, f1700])).
fof(f1700, plain, (spl93_190 <=> p304(sK78(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_190])])).
fof(f428, plain, (spl93_7 <=> ! [X7] : (p204(X7) | ~ r1(sK92, X7))), introduced(avatar_definition, [new_symbols(naming, [spl93_7])])).
fof(f1733, plain, (spl93_195 <=> p204(sK77(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_195])])).
fof(f3692, plain, (~ p304(sK78(sK92)) | (~ spl93_7 | spl93_195)), inference(subsumption_resolution, [], [f3691, f544])).
fof(f544, plain, sP6(sK92), inference(resolution, [], [f523, f248])).
fof(f248, plain, ! [X0] : (~ sP40(X0) | sP6(X0)), inference(cnf_transformation, [], [f50])).
fof(f3691, plain, (~ p304(sK78(sK92)) | ~ sP6(sK92) | (~ spl93_7 | spl93_195)), inference(subsumption_resolution, [], [f3039, f1735])).
fof(f1735, plain, (~ p204(sK77(sK92)) | spl93_195), inference(avatar_component_clause, [], [f1733])).
fof(f3039, plain, (p204(sK77(sK92)) | ~ p304(sK78(sK92)) | ~ sP6(sK92) | ~ spl93_7), inference(resolution, [], [f429, f372])).
fof(f372, plain, ! [X0] : (r1(X0, sK77(X0)) | ~ p304(sK78(X0)) | ~ sP6(X0)), inference(cnf_transformation, [], [f190])).
fof(f190, plain, ! [X0] : ((~ p204(sK77(X0)) & r1(X0, sK77(X0))) | (~ p304(sK78(X0)) & r1(X0, sK78(X0))) | ~ sP6(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK77, sK78])], [f187, f189, f188])).
fof(f188, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) => (~ p204(sK77(X0)) & r1(X0, sK77(X0)))), introduced(choice_axiom, [])).
fof(f189, plain, ! [X0] : (? [X2] : (~ p304(X2) & r1(X0, X2)) => (~ p304(sK78(X0)) & r1(X0, sK78(X0)))), introduced(choice_axiom, [])).
fof(f187, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) | ? [X2] : (~ p304(X2) & r1(X0, X2)) | ~ sP6(X0)), inference(rectify, [], [f186])).
fof(f186, plain, ! [X12] : (? [X35] : (~ p204(X35) & r1(X12, X35)) | ? [X36] : (~ p304(X36) & r1(X12, X36)) | ~ sP6(X12)), inference(nnf_transformation, [], [f13])).
fof(f429, plain, (! [X7] : (~ r1(sK92, X7) | p204(X7)) | ~ spl93_7), inference(avatar_component_clause, [], [f428])).
fof(f3690, plain, (spl93_190 | ~ spl93_12 | spl93_282), inference(avatar_split_clause, [], [f3689, f2760, f446, f1700])).
fof(f2760, plain, (spl93_282 <=> r1(sK92, sK77(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_282])])).
fof(f3689, plain, (p304(sK78(sK92)) | (~ spl93_12 | spl93_282)), inference(subsumption_resolution, [], [f3688, f544])).
fof(f3688, plain, (p304(sK78(sK92)) | ~ sP6(sK92) | (~ spl93_12 | spl93_282)), inference(subsumption_resolution, [], [f3632, f2761])).
fof(f2761, plain, (~ r1(sK92, sK77(sK92)) | spl93_282), inference(avatar_component_clause, [], [f2760])).
fof(f3632, plain, (p304(sK78(sK92)) | r1(sK92, sK77(sK92)) | ~ sP6(sK92) | ~ spl93_12), inference(resolution, [], [f447, f371])).
fof(f371, plain, ! [X0] : (r1(X0, sK78(X0)) | r1(X0, sK77(X0)) | ~ sP6(X0)), inference(cnf_transformation, [], [f190])).
fof(f3674, plain, (~ spl93_190 | ~ spl93_195), inference(avatar_split_clause, [], [f3673, f1733, f1700])).
fof(f3673, plain, (~ p304(sK78(sK92)) | ~ spl93_195), inference(subsumption_resolution, [], [f3487, f544])).
fof(f3487, plain, (~ p304(sK78(sK92)) | ~ sP6(sK92) | ~ spl93_195), inference(resolution, [], [f1734, f374])).
fof(f374, plain, ! [X0] : (~ p204(sK77(X0)) | ~ p304(sK78(X0)) | ~ sP6(X0)), inference(cnf_transformation, [], [f190])).
fof(f1734, plain, (p204(sK77(sK92)) | ~ spl93_195), inference(avatar_component_clause, [], [f1733])).
fof(f3672, plain, (spl93_190 | ~ spl93_12 | ~ spl93_195), inference(avatar_split_clause, [], [f3671, f1733, f446, f1700])).
fof(f3671, plain, (p304(sK78(sK92)) | (~ spl93_12 | ~ spl93_195)), inference(subsumption_resolution, [], [f3670, f544])).
fof(f3670, plain, (p304(sK78(sK92)) | ~ sP6(sK92) | (~ spl93_12 | ~ spl93_195)), inference(subsumption_resolution, [], [f3633, f1734])).
fof(f3633, plain, (p304(sK78(sK92)) | ~ p204(sK77(sK92)) | ~ sP6(sK92) | ~ spl93_12), inference(resolution, [], [f447, f373])).
fof(f373, plain, ! [X0] : (r1(X0, sK78(X0)) | ~ p204(sK77(X0)) | ~ sP6(X0)), inference(cnf_transformation, [], [f190])).
fof(f3583, plain, (~ spl93_257 | ~ spl93_258), inference(avatar_contradiction_clause, [], [f3582])).
fof(f3582, plain, ($false | (~ spl93_257 | ~ spl93_258)), inference(subsumption_resolution, [], [f3581, f528])).
fof(f528, plain, sP0(sK92), inference(resolution, [], [f523, f229])).
fof(f229, plain, ! [X0] : (~ sP40(X0) | sP0(X0)), inference(cnf_transformation, [], [f50])).
fof(f3581, plain, (~ sP0(sK92) | (~ spl93_257 | ~ spl93_258)), inference(subsumption_resolution, [], [f3580, f2129])).
fof(f2129, plain, (p405(sK90(sK92)) | ~ spl93_257), inference(avatar_component_clause, [], [f2128])).
fof(f2128, plain, (spl93_257 <=> p405(sK90(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_257])])).
fof(f3580, plain, (~ p405(sK90(sK92)) | ~ sP0(sK92) | ~ spl93_258), inference(resolution, [], [f2136, f398])).
fof(f398, plain, ! [X0] : (~ p305(sK89(X0)) | ~ p405(sK90(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f220])).
fof(f220, plain, ! [X0] : ((~ p305(sK89(X0)) & r1(X0, sK89(X0))) | (~ p405(sK90(X0)) & r1(X0, sK90(X0))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK89, sK90])], [f217, f219, f218])).
fof(f218, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) => (~ p305(sK89(X0)) & r1(X0, sK89(X0)))), introduced(choice_axiom, [])).
fof(f219, plain, ! [X0] : (? [X2] : (~ p405(X2) & r1(X0, X2)) => (~ p405(sK90(X0)) & r1(X0, sK90(X0)))), introduced(choice_axiom, [])).
fof(f217, plain, ! [X0] : (? [X1] : (~ p305(X1) & r1(X0, X1)) | ? [X2] : (~ p405(X2) & r1(X0, X2)) | ~ sP0(X0)), inference(rectify, [], [f216])).
fof(f216, plain, ! [X12] : (? [X57] : (~ p305(X57) & r1(X12, X57)) | ? [X58] : (~ p405(X58) & r1(X12, X58)) | ~ sP0(X12)), inference(nnf_transformation, [], [f7])).
fof(f2136, plain, (p305(sK89(sK92)) | ~ spl93_258), inference(avatar_component_clause, [], [f2134])).
fof(f2134, plain, (spl93_258 <=> p305(sK89(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_258])])).
fof(f3577, plain, (~ spl93_258 | spl93_257 | ~ spl93_16), inference(avatar_split_clause, [], [f3576, f462, f2128, f2134])).
fof(f3576, plain, (p405(sK90(sK92)) | ~ p305(sK89(sK92)) | ~ spl93_16), inference(subsumption_resolution, [], [f3241, f528])).
fof(f3241, plain, (p405(sK90(sK92)) | ~ p305(sK89(sK92)) | ~ sP0(sK92) | ~ spl93_16), inference(resolution, [], [f463, f397])).
fof(f397, plain, ! [X0] : (r1(X0, sK90(X0)) | ~ p305(sK89(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f220])).
fof(f3565, plain, (~ spl93_257 | spl93_258 | ~ spl93_11), inference(avatar_split_clause, [], [f3564, f443, f2134, f2128])).
fof(f3564, plain, (p305(sK89(sK92)) | ~ p405(sK90(sK92)) | ~ spl93_11), inference(subsumption_resolution, [], [f3441, f528])).
fof(f3441, plain, (p305(sK89(sK92)) | ~ p405(sK90(sK92)) | ~ sP0(sK92) | ~ spl93_11), inference(resolution, [], [f444, f396])).
fof(f396, plain, ! [X0] : (r1(X0, sK89(X0)) | ~ p405(sK90(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f220])).
fof(f3563, plain, (spl93_257 | ~ spl93_16 | spl93_292), inference(avatar_split_clause, [], [f3562, f2914, f462, f2128])).
fof(f2914, plain, (spl93_292 <=> r1(sK92, sK89(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_292])])).
fof(f3562, plain, (p405(sK90(sK92)) | (~ spl93_16 | spl93_292)), inference(subsumption_resolution, [], [f3561, f528])).
fof(f3561, plain, (p405(sK90(sK92)) | ~ sP0(sK92) | (~ spl93_16 | spl93_292)), inference(subsumption_resolution, [], [f3240, f2915])).
fof(f2915, plain, (~ r1(sK92, sK89(sK92)) | spl93_292), inference(avatar_component_clause, [], [f2914])).
fof(f3240, plain, (p405(sK90(sK92)) | r1(sK92, sK89(sK92)) | ~ sP0(sK92) | ~ spl93_16), inference(resolution, [], [f463, f395])).
fof(f395, plain, ! [X0] : (r1(X0, sK90(X0)) | r1(X0, sK89(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f220])).
fof(f3548, plain, (~ spl93_11 | spl93_258 | ~ spl93_292), inference(avatar_contradiction_clause, [], [f3547])).
fof(f3547, plain, ($false | (~ spl93_11 | spl93_258 | ~ spl93_292)), inference(subsumption_resolution, [], [f3544, f2135])).
fof(f2135, plain, (~ p305(sK89(sK92)) | spl93_258), inference(avatar_component_clause, [], [f2134])).
fof(f3544, plain, (p305(sK89(sK92)) | (~ spl93_11 | ~ spl93_292)), inference(resolution, [], [f2916, f444])).
fof(f2916, plain, (r1(sK92, sK89(sK92)) | ~ spl93_292), inference(avatar_component_clause, [], [f2914])).
fof(f3526, plain, (spl93_246 | ~ spl93_16 | spl93_291), inference(avatar_split_clause, [], [f3525, f2816, f462, f2056])).
fof(f3525, plain, (p405(sK88(sK92)) | (~ spl93_16 | spl93_291)), inference(subsumption_resolution, [], [f3524, f531])).
fof(f3524, plain, (p405(sK88(sK92)) | ~ sP1(sK92) | (~ spl93_16 | spl93_291)), inference(subsumption_resolution, [], [f3237, f2817])).
fof(f2817, plain, (~ r1(sK92, sK87(sK92)) | spl93_291), inference(avatar_component_clause, [], [f2816])).
fof(f3237, plain, (p405(sK88(sK92)) | r1(sK92, sK87(sK92)) | ~ sP1(sK92) | ~ spl93_16), inference(resolution, [], [f463, f391])).
fof(f391, plain, ! [X0] : (r1(X0, sK88(X0)) | r1(X0, sK87(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f215])).
fof(f3476, plain, (spl93_195 | ~ spl93_7 | ~ spl93_282), inference(avatar_split_clause, [], [f3471, f2760, f428, f1733])).
fof(f3471, plain, (p204(sK77(sK92)) | (~ spl93_7 | ~ spl93_282)), inference(resolution, [], [f2762, f429])).
fof(f2762, plain, (r1(sK92, sK77(sK92)) | ~ spl93_282), inference(avatar_component_clause, [], [f2760])).
fof(f3382, plain, (~ spl93_22 | ~ spl93_315), inference(avatar_split_clause, [], [f3381, f3342, f486])).
fof(f3342, plain, (spl93_315 <=> p204(sK58(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_315])])).
fof(f3381, plain, (~ p504(sK92) | ~ spl93_315), inference(subsumption_resolution, [], [f3352, f542])).
fof(f542, plain, sP22(sK92), inference(resolution, [], [f523, f246])).
fof(f246, plain, ! [X0] : (~ sP40(X0) | sP22(X0)), inference(cnf_transformation, [], [f50])).
fof(f3352, plain, (~ p504(sK92) | ~ sP22(sK92) | ~ spl93_315), inference(resolution, [], [f3344, f334])).
fof(f334, plain, ! [X0] : (~ p204(sK58(X0)) | ~ p504(X0) | ~ sP22(X0)), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ! [X0] : ((~ p204(sK58(X0)) & r1(X0, sK58(X0))) | ~ p504(X0) | ~ sP22(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK58])], [f120, f121])).
fof(f121, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) => (~ p204(sK58(X0)) & r1(X0, sK58(X0)))), introduced(choice_axiom, [])).
fof(f120, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) | ~ p504(X0) | ~ sP22(X0)), inference(rectify, [], [f119])).
fof(f119, plain, ! [X12] : (? [X38] : (~ p204(X38) & r1(X12, X38)) | ~ p504(X12) | ~ sP22(X12)), inference(nnf_transformation, [], [f29])).
fof(f3344, plain, (p204(sK58(sK92)) | ~ spl93_315), inference(avatar_component_clause, [], [f3342])).
fof(f3356, plain, (~ spl93_13 | ~ spl93_23), inference(avatar_split_clause, [], [f583, f490, f449])).
fof(f583, plain, (~ p503(sK92) | ~ p303(sK92)), inference(resolution, [], [f258, f523])).
fof(f258, plain, ! [X0] : (~ sP40(X0) | ~ p503(X0) | ~ p303(X0)), inference(cnf_transformation, [], [f50])).
fof(f3345, plain, (~ spl93_22 | spl93_315 | ~ spl93_7), inference(avatar_split_clause, [], [f3340, f428, f3342, f486])).
fof(f3340, plain, (p204(sK58(sK92)) | ~ p504(sK92) | ~ spl93_7), inference(subsumption_resolution, [], [f3017, f542])).
fof(f3017, plain, (p204(sK58(sK92)) | ~ p504(sK92) | ~ sP22(sK92) | ~ spl93_7), inference(resolution, [], [f429, f333])).
fof(f333, plain, ! [X0] : (r1(X0, sK58(X0)) | ~ p504(X0) | ~ sP22(X0)), inference(cnf_transformation, [], [f122])).
fof(f3261, plain, (~ spl93_16 | ~ spl93_21), inference(avatar_contradiction_clause, [], [f3260])).
fof(f3260, plain, ($false | (~ spl93_16 | ~ spl93_21)), inference(subsumption_resolution, [], [f3259, f525])).
fof(f525, plain, sP11(sK92), inference(resolution, [], [f523, f226])).
fof(f226, plain, ! [X0] : (~ sP40(X0) | sP11(X0)), inference(cnf_transformation, [], [f50])).
fof(f3259, plain, (~ sP11(sK92) | (~ spl93_16 | ~ spl93_21)), inference(subsumption_resolution, [], [f3258, f484])).
fof(f3258, plain, (~ p505(sK92) | ~ sP11(sK92) | (~ spl93_16 | ~ spl93_21)), inference(resolution, [], [f3253, f356])).
fof(f356, plain, ! [X0] : (~ p405(sK69(X0)) | ~ p505(X0) | ~ sP11(X0)), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ! [X0] : ((~ p405(sK69(X0)) & r1(X0, sK69(X0))) | ~ p505(X0) | ~ sP11(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK69])], [f164, f165])).
fof(f165, plain, ! [X0] : (? [X1] : (~ p405(X1) & r1(X0, X1)) => (~ p405(sK69(X0)) & r1(X0, sK69(X0)))), introduced(choice_axiom, [])).
fof(f164, plain, ! [X0] : (? [X1] : (~ p405(X1) & r1(X0, X1)) | ~ p505(X0) | ~ sP11(X0)), inference(rectify, [], [f163])).
fof(f163, plain, ! [X12] : (? [X61] : (~ p405(X61) & r1(X12, X61)) | ~ p505(X12) | ~ sP11(X12)), inference(nnf_transformation, [], [f18])).
fof(f3253, plain, (p405(sK69(sK92)) | (~ spl93_16 | ~ spl93_21)), inference(subsumption_resolution, [], [f3252, f525])).
fof(f3252, plain, (p405(sK69(sK92)) | ~ sP11(sK92) | (~ spl93_16 | ~ spl93_21)), inference(subsumption_resolution, [], [f3210, f484])).
fof(f3210, plain, (p405(sK69(sK92)) | ~ p505(sK92) | ~ sP11(sK92) | ~ spl93_16), inference(resolution, [], [f463, f355])).
fof(f355, plain, ! [X0] : (r1(X0, sK69(X0)) | ~ p505(X0) | ~ sP11(X0)), inference(cnf_transformation, [], [f166])).
fof(f3180, plain, (~ spl93_17 | ~ spl93_302), inference(avatar_split_clause, [], [f3179, f3165, f465])).
fof(f3165, plain, (spl93_302 <=> p204(sK57(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_302])])).
fof(f3179, plain, (~ p404(sK92) | ~ spl93_302), inference(subsumption_resolution, [], [f3175, f543])).
fof(f543, plain, sP23(sK92), inference(resolution, [], [f523, f247])).
fof(f247, plain, ! [X0] : (~ sP40(X0) | sP23(X0)), inference(cnf_transformation, [], [f50])).
fof(f3175, plain, (~ p404(sK92) | ~ sP23(sK92) | ~ spl93_302), inference(resolution, [], [f3167, f332])).
fof(f332, plain, ! [X0] : (~ p204(sK57(X0)) | ~ p404(X0) | ~ sP23(X0)), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ! [X0] : ((~ p204(sK57(X0)) & r1(X0, sK57(X0))) | ~ p404(X0) | ~ sP23(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK57])], [f116, f117])).
fof(f117, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) => (~ p204(sK57(X0)) & r1(X0, sK57(X0)))), introduced(choice_axiom, [])).
fof(f116, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) | ~ p404(X0) | ~ sP23(X0)), inference(rectify, [], [f115])).
fof(f115, plain, ! [X12] : (? [X37] : (~ p204(X37) & r1(X12, X37)) | ~ p404(X12) | ~ sP23(X12)), inference(nnf_transformation, [], [f30])).
fof(f3167, plain, (p204(sK57(sK92)) | ~ spl93_302), inference(avatar_component_clause, [], [f3165])).
fof(f3168, plain, (~ spl93_17 | spl93_302 | ~ spl93_7), inference(avatar_split_clause, [], [f3163, f428, f3165, f465])).
fof(f3163, plain, (p204(sK57(sK92)) | ~ p404(sK92) | ~ spl93_7), inference(subsumption_resolution, [], [f3016, f543])).
fof(f3016, plain, (p204(sK57(sK92)) | ~ p404(sK92) | ~ sP23(sK92) | ~ spl93_7), inference(resolution, [], [f429, f331])).
fof(f331, plain, ! [X0] : (r1(X0, sK57(X0)) | ~ p404(X0) | ~ sP23(X0)), inference(cnf_transformation, [], [f118])).
fof(f3150, plain, (~ spl93_27 | ~ spl93_297), inference(avatar_split_clause, [], [f3149, f3127, f507])).
fof(f3127, plain, (spl93_297 <=> p204(sK59(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_297])])).
fof(f3149, plain, (~ p604(sK92) | ~ spl93_297), inference(subsumption_resolution, [], [f3137, f541])).
fof(f541, plain, sP21(sK92), inference(resolution, [], [f523, f245])).
fof(f245, plain, ! [X0] : (~ sP40(X0) | sP21(X0)), inference(cnf_transformation, [], [f50])).
fof(f3137, plain, (~ p604(sK92) | ~ sP21(sK92) | ~ spl93_297), inference(resolution, [], [f3129, f336])).
fof(f336, plain, ! [X0] : (~ p204(sK59(X0)) | ~ p604(X0) | ~ sP21(X0)), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ! [X0] : ((~ p204(sK59(X0)) & r1(X0, sK59(X0))) | ~ p604(X0) | ~ sP21(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK59])], [f124, f125])).
fof(f125, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) => (~ p204(sK59(X0)) & r1(X0, sK59(X0)))), introduced(choice_axiom, [])).
fof(f124, plain, ! [X0] : (? [X1] : (~ p204(X1) & r1(X0, X1)) | ~ p604(X0) | ~ sP21(X0)), inference(rectify, [], [f123])).
fof(f123, plain, ! [X12] : (? [X39] : (~ p204(X39) & r1(X12, X39)) | ~ p604(X12) | ~ sP21(X12)), inference(nnf_transformation, [], [f28])).
fof(f3129, plain, (p204(sK59(sK92)) | ~ spl93_297), inference(avatar_component_clause, [], [f3127])).
fof(f3130, plain, (~ spl93_27 | spl93_297 | ~ spl93_7), inference(avatar_split_clause, [], [f3125, f428, f3127, f507])).
fof(f3125, plain, (p204(sK59(sK92)) | ~ p604(sK92) | ~ spl93_7), inference(subsumption_resolution, [], [f3018, f541])).
fof(f3018, plain, (p204(sK59(sK92)) | ~ p604(sK92) | ~ sP21(sK92) | ~ spl93_7), inference(resolution, [], [f429, f335])).
fof(f335, plain, ! [X0] : (r1(X0, sK59(X0)) | ~ p604(X0) | ~ sP21(X0)), inference(cnf_transformation, [], [f126])).
fof(f3118, plain, (spl93_278 | spl93_166 | ~ spl93_7), inference(avatar_split_clause, [], [f3117, f428, f1550, f2736])).
fof(f2736, plain, (spl93_278 <=> r1(sK92, sK73(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_278])])).
fof(f1550, plain, (spl93_166 <=> p204(sK74(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_166])])).
fof(f3117, plain, (p204(sK74(sK92)) | r1(sK92, sK73(sK92)) | ~ spl93_7), inference(subsumption_resolution, [], [f3034, f549])).
fof(f549, plain, sP8(sK92), inference(resolution, [], [f523, f253])).
fof(f253, plain, ! [X0] : (~ sP40(X0) | sP8(X0)), inference(cnf_transformation, [], [f50])).
fof(f3034, plain, (p204(sK74(sK92)) | r1(sK92, sK73(sK92)) | ~ sP8(sK92) | ~ spl93_7), inference(resolution, [], [f429, f363])).
fof(f363, plain, ! [X0] : (r1(X0, sK74(X0)) | r1(X0, sK73(X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ! [X0] : ((~ p104(sK73(X0)) & r1(X0, sK73(X0))) | (~ p204(sK74(X0)) & r1(X0, sK74(X0))) | ~ sP8(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK73, sK74])], [f177, f179, f178])).
fof(f178, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) => (~ p104(sK73(X0)) & r1(X0, sK73(X0)))), introduced(choice_axiom, [])).
fof(f179, plain, ! [X0] : (? [X2] : (~ p204(X2) & r1(X0, X2)) => (~ p204(sK74(X0)) & r1(X0, sK74(X0)))), introduced(choice_axiom, [])).
fof(f177, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) | ? [X2] : (~ p204(X2) & r1(X0, X2)) | ~ sP8(X0)), inference(rectify, [], [f176])).
fof(f176, plain, ! [X12] : (? [X28] : (~ p104(X28) & r1(X12, X28)) | ? [X29] : (~ p204(X29) & r1(X12, X29)) | ~ sP8(X12)), inference(nnf_transformation, [], [f15])).
fof(f3110, plain, (~ spl93_169 | spl93_166 | ~ spl93_7), inference(avatar_split_clause, [], [f3109, f428, f1550, f1568])).
fof(f1568, plain, (spl93_169 <=> p104(sK73(sK92))), introduced(avatar_definition, [new_symbols(naming, [spl93_169])])).
fof(f3109, plain, (p204(sK74(sK92)) | ~ p104(sK73(sK92)) | ~ spl93_7), inference(subsumption_resolution, [], [f3035, f549])).
fof(f3035, plain, (p204(sK74(sK92)) | ~ p104(sK73(sK92)) | ~ sP8(sK92) | ~ spl93_7), inference(resolution, [], [f429, f365])).
fof(f365, plain, ! [X0] : (r1(X0, sK74(X0)) | ~ p104(sK73(X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f180])).
fof(f3108, plain, (~ spl93_166 | spl93_169 | ~ spl93_2), inference(avatar_split_clause, [], [f3107, f411, f1568, f1550])).
fof(f3107, plain, (p104(sK73(sK92)) | ~ p204(sK74(sK92)) | ~ spl93_2), inference(subsumption_resolution, [], [f2955, f549])).
fof(f2955, plain, (p104(sK73(sK92)) | ~ p204(sK74(sK92)) | ~ sP8(sK92) | ~ spl93_2), inference(resolution, [], [f412, f364])).
fof(f364, plain, ! [X0] : (r1(X0, sK73(X0)) | ~ p204(sK74(X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f180])).
fof(f3106, plain, (~ spl93_166 | ~ spl93_169), inference(avatar_contradiction_clause, [], [f3105])).
fof(f3105, plain, ($false | (~ spl93_166 | ~ spl93_169)), inference(subsumption_resolution, [], [f3104, f549])).
fof(f3104, plain, (~ sP8(sK92) | (~ spl93_166 | ~ spl93_169)), inference(subsumption_resolution, [], [f3103, f1551])).
fof(f1551, plain, (p204(sK74(sK92)) | ~ spl93_166), inference(avatar_component_clause, [], [f1550])).
fof(f3103, plain, (~ p204(sK74(sK92)) | ~ sP8(sK92) | ~ spl93_169), inference(resolution, [], [f1570, f366])).
fof(f366, plain, ! [X0] : (~ p104(sK73(X0)) | ~ p204(sK74(X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f180])).
fof(f1570, plain, (p104(sK73(sK92)) | ~ spl93_169), inference(avatar_component_clause, [], [f1568])).
fof(f3094, plain, (spl93_169 | ~ spl93_2 | ~ spl93_278), inference(avatar_split_clause, [], [f3091, f2736, f411, f1568])).
fof(f3091, plain, (p104(sK73(sK92)) | (~ spl93_2 | ~ spl93_278)), inference(resolution, [], [f2738, f412])).
fof(f2738, plain, (r1(sK92, sK73(sK92)) | ~ spl93_278), inference(avatar_component_clause, [], [f2736])).
fof(f2921, plain, (~ spl93_6 | ~ spl93_21), inference(avatar_contradiction_clause, [], [f2920])).
fof(f2920, plain, ($false | (~ spl93_6 | ~ spl93_21)), inference(subsumption_resolution, [], [f2919, f530])).
fof(f530, plain, sP15(sK92), inference(resolution, [], [f523, f231])).
fof(f231, plain, ! [X0] : (~ sP40(X0) | sP15(X0)), inference(cnf_transformation, [], [f50])).
fof(f2919, plain, (~ sP15(sK92) | (~ spl93_6 | ~ spl93_21)), inference(subsumption_resolution, [], [f2918, f484])).
fof(f2918, plain, (~ p505(sK92) | ~ sP15(sK92) | (~ spl93_6 | ~ spl93_21)), inference(resolution, [], [f2892, f348])).
fof(f348, plain, ! [X0] : (~ p205(sK65(X0)) | ~ p505(X0) | ~ sP15(X0)), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ! [X0] : ((~ p205(sK65(X0)) & r1(X0, sK65(X0))) | ~ p505(X0) | ~ sP15(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK65])], [f148, f149])).
fof(f149, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) => (~ p205(sK65(X0)) & r1(X0, sK65(X0)))), introduced(choice_axiom, [])).
fof(f148, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) | ~ p505(X0) | ~ sP15(X0)), inference(rectify, [], [f147])).
fof(f147, plain, ! [X12] : (? [X55] : (~ p205(X55) & r1(X12, X55)) | ~ p505(X12) | ~ sP15(X12)), inference(nnf_transformation, [], [f22])).
fof(f2892, plain, (p205(sK65(sK92)) | (~ spl93_6 | ~ spl93_21)), inference(subsumption_resolution, [], [f2891, f530])).
fof(f2891, plain, (p205(sK65(sK92)) | ~ sP15(sK92) | (~ spl93_6 | ~ spl93_21)), inference(subsumption_resolution, [], [f2848, f484])).
fof(f2848, plain, (p205(sK65(sK92)) | ~ p505(sK92) | ~ sP15(sK92) | ~ spl93_6), inference(resolution, [], [f426, f347])).
fof(f347, plain, ! [X0] : (r1(X0, sK65(X0)) | ~ p505(X0) | ~ sP15(X0)), inference(cnf_transformation, [], [f150])).
fof(f2907, plain, (spl93_284 | spl93_203 | ~ spl93_6), inference(avatar_split_clause, [], [f2906, f425, f1779, f2772])).
fof(f2906, plain, (p205(sK80(sK92)) | r1(sK92, sK79(sK92)) | ~ spl93_6), inference(subsumption_resolution, [], [f2867, f537])).
fof(f2867, plain, (p205(sK80(sK92)) | r1(sK92, sK79(sK92)) | ~ sP5(sK92) | ~ spl93_6), inference(resolution, [], [f426, f375])).
fof(f375, plain, ! [X0] : (r1(X0, sK80(X0)) | r1(X0, sK79(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f195])).
fof(f2823, plain, (~ spl93_8 | ~ spl93_13), inference(avatar_contradiction_clause, [], [f2822])).
fof(f2822, plain, ($false | (~ spl93_8 | ~ spl93_13)), inference(subsumption_resolution, [], [f2821, f553])).
fof(f553, plain, sP30(sK92), inference(resolution, [], [f523, f263])).
fof(f263, plain, ! [X0] : (~ sP40(X0) | sP30(X0)), inference(cnf_transformation, [], [f50])).
fof(f2821, plain, (~ sP30(sK92) | (~ spl93_8 | ~ spl93_13)), inference(subsumption_resolution, [], [f2820, f451])).
fof(f2820, plain, (~ p303(sK92) | ~ sP30(sK92) | (~ spl93_8 | ~ spl93_13)), inference(resolution, [], [f2711, f318])).
fof(f318, plain, ! [X0] : (~ p203(sK50(X0)) | ~ p303(X0) | ~ sP30(X0)), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ! [X0] : ((~ p203(sK50(X0)) & r1(X0, sK50(X0))) | ~ p303(X0) | ~ sP30(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK50])], [f88, f89])).
fof(f89, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) => (~ p203(sK50(X0)) & r1(X0, sK50(X0)))), introduced(choice_axiom, [])).
fof(f88, plain, ! [X0] : (? [X1] : (~ p203(X1) & r1(X0, X1)) | ~ p303(X0) | ~ sP30(X0)), inference(rectify, [], [f87])).
fof(f87, plain, ! [X12] : (? [X24] : (~ p203(X24) & r1(X12, X24)) | ~ p303(X12) | ~ sP30(X12)), inference(nnf_transformation, [], [f37])).
fof(f2711, plain, (p203(sK50(sK92)) | (~ spl93_8 | ~ spl93_13)), inference(subsumption_resolution, [], [f2710, f553])).
fof(f2710, plain, (p203(sK50(sK92)) | ~ sP30(sK92) | (~ spl93_8 | ~ spl93_13)), inference(subsumption_resolution, [], [f2664, f451])).
fof(f2664, plain, (p203(sK50(sK92)) | ~ p303(sK92) | ~ sP30(sK92) | ~ spl93_8), inference(resolution, [], [f432, f317])).
fof(f317, plain, ! [X0] : (r1(X0, sK50(X0)) | ~ p303(X0) | ~ sP30(X0)), inference(cnf_transformation, [], [f90])).
fof(f2654, plain, (~ spl93_18 | ~ spl93_13), inference(avatar_split_clause, [], [f2653, f449, f469])).
fof(f2653, plain, (~ p403(sK92) | ~ spl93_13), inference(subsumption_resolution, [], [f584, f451])).
fof(f584, plain, (~ p403(sK92) | ~ p303(sK92)), inference(resolution, [], [f259, f523])).
fof(f259, plain, ! [X0] : (~ sP40(X0) | ~ p403(X0) | ~ p303(X0)), inference(cnf_transformation, [], [f50])).
fof(f2648, plain, (~ spl93_14 | ~ spl93_9), inference(avatar_split_clause, [], [f2647, f434, f453])).
fof(f2647, plain, (~ p302(sK92) | ~ spl93_9), inference(subsumption_resolution, [], [f598, f436])).
fof(f598, plain, (~ p302(sK92) | ~ p202(sK92)), inference(resolution, [], [f278, f523])).
fof(f278, plain, ! [X0] : (~ sP40(X0) | ~ p302(X0) | ~ p202(X0)), inference(cnf_transformation, [], [f50])).
fof(f2563, plain, (~ spl93_9 | ~ spl93_19), inference(avatar_split_clause, [], [f597, f473, f434])).
fof(f597, plain, (~ p402(sK92) | ~ p202(sK92)), inference(resolution, [], [f277, f523])).
fof(f277, plain, ! [X0] : (~ sP40(X0) | ~ p402(X0) | ~ p202(X0)), inference(cnf_transformation, [], [f50])).
fof(f2562, plain, (~ spl93_10 | ~ spl93_5), inference(avatar_split_clause, [], [f2561, f420, f438])).
fof(f2561, plain, (~ p201(sK92) | ~ spl93_5), inference(subsumption_resolution, [], [f618, f422])).
fof(f422, plain, (p101(sK92) | ~ spl93_5), inference(avatar_component_clause, [], [f420])).
fof(f618, plain, (~ p201(sK92) | ~ p101(sK92)), inference(resolution, [], [f298, f523])).
fof(f298, plain, ! [X0] : (~ sP40(X0) | ~ p201(X0) | ~ p101(X0)), inference(cnf_transformation, [], [f50])).
fof(f2552, plain, (spl93_264 | ~ spl93_8 | spl93_153), inference(avatar_split_clause, [], [f2551, f1471, f431, f2183])).
fof(f2551, plain, (r1(sK92, sK71(sK92)) | (~ spl93_8 | spl93_153)), inference(subsumption_resolution, [], [f2550, f558])).
fof(f2550, plain, (r1(sK92, sK71(sK92)) | ~ sP9(sK92) | (~ spl93_8 | spl93_153)), inference(subsumption_resolution, [], [f2516, f1473])).
fof(f1473, plain, (~ p203(sK72(sK92)) | spl93_153), inference(avatar_component_clause, [], [f1471])).
fof(f2516, plain, (p203(sK72(sK92)) | r1(sK92, sK71(sK92)) | ~ sP9(sK92) | ~ spl93_8), inference(resolution, [], [f432, f359])).
fof(f359, plain, ! [X0] : (r1(X0, sK72(X0)) | r1(X0, sK71(X0)) | ~ sP9(X0)), inference(cnf_transformation, [], [f175])).
fof(f2019, plain, (~ spl93_237 | ~ spl93_11 | spl93_235), inference(avatar_split_clause, [], [f2018, f1984, f443, f1996])).
fof(f2018, plain, (~ p205(sK85(sK92)) | (~ spl93_11 | spl93_235)), inference(subsumption_resolution, [], [f2017, f1986])).
fof(f1986, plain, (~ p305(sK86(sK92)) | spl93_235), inference(avatar_component_clause, [], [f1984])).
fof(f2017, plain, (~ p205(sK85(sK92)) | p305(sK86(sK92)) | ~ spl93_11), inference(subsumption_resolution, [], [f2008, f532])).
fof(f2008, plain, (~ p205(sK85(sK92)) | ~ sP2(sK92) | p305(sK86(sK92)) | ~ spl93_11), inference(resolution, [], [f389, f444])).
fof(f389, plain, ! [X0] : (r1(X0, sK86(X0)) | ~ p205(sK85(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f210])).
fof(f1999, plain, (spl93_237 | ~ spl93_235 | ~ spl93_6), inference(avatar_split_clause, [], [f1994, f425, f1984, f1996])).
fof(f1994, plain, (~ p305(sK86(sK92)) | p205(sK85(sK92)) | ~ spl93_6), inference(subsumption_resolution, [], [f1963, f532])).
fof(f1963, plain, (~ p305(sK86(sK92)) | ~ sP2(sK92) | p205(sK85(sK92)) | ~ spl93_6), inference(resolution, [], [f388, f426])).
fof(f388, plain, ! [X0] : (r1(X0, sK85(X0)) | ~ p305(sK86(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f210])).
fof(f1824, plain, (~ spl93_208 | ~ spl93_6 | spl93_203), inference(avatar_split_clause, [], [f1823, f1779, f425, f1812])).
fof(f1823, plain, (~ p105(sK79(sK92)) | (~ spl93_6 | spl93_203)), inference(subsumption_resolution, [], [f1822, f1781])).
fof(f1781, plain, (~ p205(sK80(sK92)) | spl93_203), inference(avatar_component_clause, [], [f1779])).
fof(f1822, plain, (~ p105(sK79(sK92)) | p205(sK80(sK92)) | ~ spl93_6), inference(subsumption_resolution, [], [f1804, f537])).
fof(f1804, plain, (~ p105(sK79(sK92)) | ~ sP5(sK92) | p205(sK80(sK92)) | ~ spl93_6), inference(resolution, [], [f377, f426])).
fof(f377, plain, ! [X0] : (r1(X0, sK80(X0)) | ~ p105(sK79(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f195])).
fof(f1428, plain, (spl93_145 | ~ spl93_27 | ~ spl93_2), inference(avatar_split_clause, [], [f1423, f411, f507, f1425])).
fof(f1423, plain, (~ p604(sK92) | p104(sK56(sK92)) | ~ spl93_2), inference(subsumption_resolution, [], [f924, f545])).
fof(f924, plain, (~ p604(sK92) | ~ sP24(sK92) | p104(sK56(sK92)) | ~ spl93_2), inference(resolution, [], [f329, f412])).
fof(f329, plain, ! [X0] : (r1(X0, sK56(X0)) | ~ p604(X0) | ~ sP24(X0)), inference(cnf_transformation, [], [f114])).
fof(f1278, plain, (~ spl93_6 | ~ spl93_26), inference(avatar_contradiction_clause, [], [f1277])).
fof(f1277, plain, ($false | (~ spl93_6 | ~ spl93_26)), inference(subsumption_resolution, [], [f1276, f529])).
fof(f529, plain, sP14(sK92), inference(resolution, [], [f523, f230])).
fof(f230, plain, ! [X0] : (~ sP40(X0) | sP14(X0)), inference(cnf_transformation, [], [f50])).
fof(f1276, plain, (~ sP14(sK92) | (~ spl93_6 | ~ spl93_26)), inference(subsumption_resolution, [], [f1275, f505])).
fof(f1275, plain, (~ p605(sK92) | ~ sP14(sK92) | (~ spl93_6 | ~ spl93_26)), inference(resolution, [], [f1088, f350])).
fof(f350, plain, ! [X0] : (~ p205(sK66(X0)) | ~ p605(X0) | ~ sP14(X0)), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ! [X0] : ((~ p205(sK66(X0)) & r1(X0, sK66(X0))) | ~ p605(X0) | ~ sP14(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK66])], [f152, f153])).
fof(f153, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) => (~ p205(sK66(X0)) & r1(X0, sK66(X0)))), introduced(choice_axiom, [])).
fof(f152, plain, ! [X0] : (? [X1] : (~ p205(X1) & r1(X0, X1)) | ~ p605(X0) | ~ sP14(X0)), inference(rectify, [], [f151])).
fof(f151, plain, ! [X12] : (? [X56] : (~ p205(X56) & r1(X12, X56)) | ~ p605(X12) | ~ sP14(X12)), inference(nnf_transformation, [], [f21])).
fof(f1088, plain, (p205(sK66(sK92)) | (~ spl93_6 | ~ spl93_26)), inference(subsumption_resolution, [], [f1087, f529])).
fof(f1087, plain, (~ sP14(sK92) | p205(sK66(sK92)) | (~ spl93_6 | ~ spl93_26)), inference(subsumption_resolution, [], [f1072, f505])).
fof(f1072, plain, (~ p605(sK92) | ~ sP14(sK92) | p205(sK66(sK92)) | ~ spl93_6), inference(resolution, [], [f349, f426])).
fof(f349, plain, ! [X0] : (r1(X0, sK66(X0)) | ~ p605(X0) | ~ sP14(X0)), inference(cnf_transformation, [], [f154])).
fof(f1117, plain, (~ spl93_2 | ~ spl93_22), inference(avatar_contradiction_clause, [], [f1116])).
fof(f1116, plain, ($false | (~ spl93_2 | ~ spl93_22)), inference(subsumption_resolution, [], [f1115, f546])).
fof(f546, plain, sP25(sK92), inference(resolution, [], [f523, f250])).
fof(f250, plain, ! [X0] : (~ sP40(X0) | sP25(X0)), inference(cnf_transformation, [], [f50])).
fof(f1115, plain, (~ sP25(sK92) | (~ spl93_2 | ~ spl93_22)), inference(subsumption_resolution, [], [f1114, f488])).
fof(f488, plain, (p504(sK92) | ~ spl93_22), inference(avatar_component_clause, [], [f486])).
fof(f1114, plain, (~ p504(sK92) | ~ sP25(sK92) | (~ spl93_2 | ~ spl93_22)), inference(resolution, [], [f919, f328])).
fof(f328, plain, ! [X0] : (~ p104(sK55(X0)) | ~ p504(X0) | ~ sP25(X0)), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ! [X0] : ((~ p104(sK55(X0)) & r1(X0, sK55(X0))) | ~ p504(X0) | ~ sP25(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK55])], [f108, f109])).
fof(f109, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) => (~ p104(sK55(X0)) & r1(X0, sK55(X0)))), introduced(choice_axiom, [])).
fof(f108, plain, ! [X0] : (? [X1] : (~ p104(X1) & r1(X0, X1)) | ~ p504(X0) | ~ sP25(X0)), inference(rectify, [], [f107])).
fof(f107, plain, ! [X12] : (? [X33] : (~ p104(X33) & r1(X12, X33)) | ~ p504(X12) | ~ sP25(X12)), inference(nnf_transformation, [], [f32])).
fof(f919, plain, (p104(sK55(sK92)) | (~ spl93_2 | ~ spl93_22)), inference(subsumption_resolution, [], [f918, f546])).
fof(f918, plain, (~ sP25(sK92) | p104(sK55(sK92)) | (~ spl93_2 | ~ spl93_22)), inference(subsumption_resolution, [], [f898, f488])).
fof(f898, plain, (~ p504(sK92) | ~ sP25(sK92) | p104(sK55(sK92)) | ~ spl93_2), inference(resolution, [], [f327, f412])).
fof(f327, plain, ! [X0] : (r1(X0, sK55(X0)) | ~ p504(X0) | ~ sP25(X0)), inference(cnf_transformation, [], [f110])).
fof(f614, plain, (~ spl93_5 | ~ spl93_30), inference(avatar_split_clause, [], [f613, f519, f420])).
fof(f519, plain, (spl93_30 <=> p601(sK92)), introduced(avatar_definition, [new_symbols(naming, [spl93_30])])).
fof(f613, plain, (~ p601(sK92) | ~ p101(sK92)), inference(resolution, [], [f294, f523])).
fof(f294, plain, ! [X0] : (~ sP40(X0) | ~ p601(X0) | ~ p101(X0)), inference(cnf_transformation, [], [f50])).
fof(f609, plain, (~ spl93_10 | ~ spl93_30), inference(avatar_split_clause, [], [f608, f519, f438])).
fof(f608, plain, (~ p601(sK92) | ~ p201(sK92)), inference(resolution, [], [f290, f523])).
fof(f290, plain, ! [X0] : (~ sP40(X0) | ~ p601(X0) | ~ p201(X0)), inference(cnf_transformation, [], [f50])).
fof(f605, plain, (~ spl93_15 | ~ spl93_30), inference(avatar_split_clause, [], [f604, f519, f457])).
fof(f604, plain, (~ p601(sK92) | ~ p301(sK92)), inference(resolution, [], [f287, f523])).
fof(f287, plain, ! [X0] : (~ sP40(X0) | ~ p601(X0) | ~ p301(X0)), inference(cnf_transformation, [], [f50])).
fof(f602, plain, (~ spl93_20 | ~ spl93_30), inference(avatar_split_clause, [], [f601, f519, f477])).
fof(f601, plain, (~ p601(sK92) | ~ p401(sK92)), inference(resolution, [], [f285, f523])).
fof(f285, plain, ! [X0] : (~ sP40(X0) | ~ p601(X0) | ~ p401(X0)), inference(cnf_transformation, [], [f50])).
fof(f600, plain, (~ spl93_25 | ~ spl93_30), inference(avatar_split_clause, [], [f599, f519, f498])).
fof(f599, plain, (~ p601(sK92) | ~ p501(sK92)), inference(resolution, [], [f284, f523])).
fof(f284, plain, ! [X0] : (~ sP40(X0) | ~ p601(X0) | ~ p501(X0)), inference(cnf_transformation, [], [f50])).
fof(f595, plain, (~ spl93_9 | ~ spl93_29), inference(avatar_split_clause, [], [f594, f515, f434])).
fof(f594, plain, (~ p602(sK92) | ~ p202(sK92)), inference(resolution, [], [f275, f523])).
fof(f275, plain, ! [X0] : (~ sP40(X0) | ~ p602(X0) | ~ p202(X0)), inference(cnf_transformation, [], [f50])).
fof(f591, plain, (~ spl93_14 | ~ spl93_29), inference(avatar_split_clause, [], [f590, f515, f453])).
fof(f590, plain, (~ p602(sK92) | ~ p302(sK92)), inference(resolution, [], [f272, f523])).
fof(f272, plain, ! [X0] : (~ sP40(X0) | ~ p602(X0) | ~ p302(X0)), inference(cnf_transformation, [], [f50])).
fof(f588, plain, (~ spl93_19 | ~ spl93_29), inference(avatar_split_clause, [], [f587, f515, f473])).
fof(f587, plain, (~ p602(sK92) | ~ p402(sK92)), inference(resolution, [], [f270, f523])).
fof(f270, plain, ! [X0] : (~ sP40(X0) | ~ p602(X0) | ~ p402(X0)), inference(cnf_transformation, [], [f50])).
fof(f586, plain, (~ spl93_24 | ~ spl93_29), inference(avatar_split_clause, [], [f585, f515, f494])).
fof(f585, plain, (~ p602(sK92) | ~ p502(sK92)), inference(resolution, [], [f269, f523])).
fof(f269, plain, ! [X0] : (~ sP40(X0) | ~ p602(X0) | ~ p502(X0)), inference(cnf_transformation, [], [f50])).
fof(f582, plain, (~ spl93_13 | ~ spl93_28), inference(avatar_split_clause, [], [f581, f511, f449])).
fof(f581, plain, (~ p603(sK92) | ~ p303(sK92)), inference(resolution, [], [f257, f523])).
fof(f257, plain, ! [X0] : (~ sP40(X0) | ~ p603(X0) | ~ p303(X0)), inference(cnf_transformation, [], [f50])).
fof(f579, plain, (~ spl93_18 | ~ spl93_28), inference(avatar_split_clause, [], [f578, f511, f469])).
fof(f578, plain, (~ p603(sK92) | ~ p403(sK92)), inference(resolution, [], [f255, f523])).
fof(f255, plain, ! [X0] : (~ sP40(X0) | ~ p603(X0) | ~ p403(X0)), inference(cnf_transformation, [], [f50])).
fof(f577, plain, (~ spl93_23 | ~ spl93_28), inference(avatar_split_clause, [], [f576, f511, f490])).
fof(f576, plain, (~ p603(sK92) | ~ p503(sK92)), inference(resolution, [], [f254, f523])).
fof(f254, plain, ! [X0] : (~ sP40(X0) | ~ p603(X0) | ~ p503(X0)), inference(cnf_transformation, [], [f50])).
fof(f522, plain, (spl93_26 | spl93_27 | spl93_28 | spl93_29 | spl93_30), inference(avatar_split_clause, [], [f401, f519, f515, f511, f507, f503])).
fof(f401, plain, (p601(sK92) | p602(sK92) | p603(sK92) | p604(sK92) | p605(sK92)), inference(cnf_transformation, [], [f223])).
fof(f501, plain, (spl93_21 | spl93_22 | spl93_23 | spl93_24 | spl93_25), inference(avatar_split_clause, [], [f402, f498, f494, f490, f486, f482])).
fof(f402, plain, (p501(sK92) | p502(sK92) | p503(sK92) | p504(sK92) | p505(sK92)), inference(cnf_transformation, [], [f223])).
fof(f480, plain, (spl93_16 | spl93_17 | spl93_18 | spl93_19 | spl93_20), inference(avatar_split_clause, [], [f403, f477, f473, f469, f465, f462])).
fof(f403, plain, ! [X11] : (p401(sK92) | p402(sK92) | p403(sK92) | p404(sK92) | p405(X11) | ~ r1(sK92, X11)), inference(cnf_transformation, [], [f223])).
fof(f460, plain, (spl93_11 | spl93_12 | spl93_13 | spl93_14 | spl93_15), inference(avatar_split_clause, [], [f404, f457, f453, f449, f446, f443])).
fof(f404, plain, ! [X10, X9] : (p301(sK92) | p302(sK92) | p303(sK92) | p304(X9) | ~ r1(sK92, X9) | p305(X10) | ~ r1(sK92, X10)), inference(cnf_transformation, [], [f223])).
fof(f441, plain, (spl93_6 | spl93_7 | spl93_8 | spl93_9 | spl93_10), inference(avatar_split_clause, [], [f405, f438, f434, f431, f428, f425])).
fof(f405, plain, ! [X6, X8, X7] : (p201(sK92) | p202(sK92) | p203(X6) | ~ r1(sK92, X6) | p204(X7) | ~ r1(sK92, X7) | p205(X8) | ~ r1(sK92, X8)), inference(cnf_transformation, [], [f223])).
fof(f423, plain, (spl93_1 | spl93_2 | spl93_3 | spl93_4 | spl93_5), inference(avatar_split_clause, [], [f406, f420, f417, f414, f411, f408])).
fof(f406, plain, ! [X4, X2, X5, X3] : (p101(sK92) | p102(X2) | ~ r1(sK92, X2) | p103(X3) | ~ r1(sK92, X3) | p104(X4) | ~ r1(sK92, X4) | p105(X5) | ~ r1(sK92, X5)), inference(cnf_transformation, [], [f223])).