fof(f373175, plain, $false, inference(avatar_sat_refutation, [], [f2913, f2920, f2933, f2959, f2983, f3007, f3067, f5433, f5537, f5734, f14255, f14328, f14481, f30923, f31537, f31730, f37861, f62018, f62229, f72028, f112709, f112907, f373174])).
fof(f373174, plain, (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815 | ~ spl63_11303 | ~ spl63_11306 | spl63_11307), inference(avatar_contradiction_clause, [], [f373173])).
fof(f373173, plain, ($false | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815 | ~ spl63_11303 | ~ spl63_11306 | spl63_11307)), inference(subsumption_resolution, [], [f373172, f67761])).
fof(f67761, plain, (sP8(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ spl63_11306), inference(avatar_component_clause, [], [f67760])).
fof(f67760, plain, (spl63_11306 <=> sP8(sK42(sK40(sK38(sK36(sK35(sK62))))))), introduced(avatar_definition, [new_symbols(naming, [spl63_11306])])).
fof(f373172, plain, (~ sP8(sK42(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815 | ~ spl63_11303 | ~ spl63_11306 | spl63_11307)), inference(subsumption_resolution, [], [f373171, f67750])).
fof(f67750, plain, (p105(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ spl63_11303), inference(avatar_component_clause, [], [f67748])).
fof(f67748, plain, (spl63_11303 <=> p105(sK42(sK40(sK38(sK36(sK35(sK62))))))), introduced(avatar_definition, [new_symbols(naming, [spl63_11303])])).
fof(f373171, plain, (~ p105(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ sP8(sK42(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815 | ~ spl63_11303 | ~ spl63_11306 | spl63_11307)), inference(subsumption_resolution, [], [f373170, f67765])).
fof(f67765, plain, (~ p106(sK42(sK40(sK38(sK36(sK35(sK62)))))) | spl63_11307), inference(avatar_component_clause, [], [f67764])).
fof(f67764, plain, (spl63_11307 <=> p106(sK42(sK40(sK38(sK36(sK35(sK62))))))), introduced(avatar_definition, [new_symbols(naming, [spl63_11307])])).
fof(f373170, plain, (p106(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ p105(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ sP8(sK42(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815 | ~ spl63_11303 | ~ spl63_11306 | spl63_11307)), inference(resolution, [], [f373138, f280])).
fof(f280, plain, ! [X0] : (~ p7(sK45(X0)) | p106(X0) | ~ p105(X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ! [X0] : (~ p105(X0) | p106(X0) | ((p106(sK44(X0)) & ~ p107(sK44(X0)) & p7(sK44(X0)) & r1(X0, sK44(X0))) & (p106(sK45(X0)) & ~ p107(sK45(X0)) & ~ p7(sK45(X0)) & r1(X0, sK45(X0)))) | ~ sP8(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK44, sK45])], [f107, f109, f108])).
fof(f108, plain, ! [X0] : (? [X1] : (p106(X1) & ~ p107(X1) & p7(X1) & r1(X0, X1)) => (p106(sK44(X0)) & ~ p107(sK44(X0)) & p7(sK44(X0)) & r1(X0, sK44(X0)))), introduced(choice_axiom, [])).
fof(f109, plain, ! [X0] : (? [X2] : (p106(X2) & ~ p107(X2) & ~ p7(X2) & r1(X0, X2)) => (p106(sK45(X0)) & ~ p107(sK45(X0)) & ~ p7(sK45(X0)) & r1(X0, sK45(X0)))), introduced(choice_axiom, [])).
fof(f107, plain, ! [X0] : (~ p105(X0) | p106(X0) | (? [X1] : (p106(X1) & ~ p107(X1) & p7(X1) & r1(X0, X1)) & ? [X2] : (p106(X2) & ~ p107(X2) & ~ p7(X2) & r1(X0, X2))) | ~ sP8(X0)), inference(rectify, [], [f106])).
fof(f106, plain, ! [X15] : (~ p105(X15) | p106(X15) | (? [X58] : (p106(X58) & ~ p107(X58) & p7(X58) & r1(X15, X58)) & ? [X59] : (p106(X59) & ~ p107(X59) & ~ p7(X59) & r1(X15, X59))) | ~ sP8(X15)), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ! [X15] : (~ p105(X15) | p106(X15) | (? [X58] : (p106(X58) & ~ p107(X58) & p7(X58) & r1(X15, X58)) & ? [X59] : (p106(X59) & ~ p107(X59) & ~ p7(X59) & r1(X15, X59))) | ~ sP8(X15)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X15] : (sP8(X15) <=> (~ p105(X15) | p106(X15) | (? [X58] : (p106(X58) & ~ p107(X58) & p7(X58) & r1(X15, X58)) & ? [X59] : (p106(X59) & ~ p107(X59) & ~ p7(X59) & r1(X15, X59))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f373138, plain, (p7(sK45(sK42(sK40(sK38(sK36(sK35(sK62))))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815 | ~ spl63_11303 | ~ spl63_11306 | spl63_11307)), inference(subsumption_resolution, [], [f373137, f67761])).
fof(f373137, plain, (p7(sK45(sK42(sK40(sK38(sK36(sK35(sK62))))))) | ~ sP8(sK42(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815 | ~ spl63_11303 | spl63_11307)), inference(subsumption_resolution, [], [f373136, f67750])).
fof(f373136, plain, (p7(sK45(sK42(sK40(sK38(sK36(sK35(sK62))))))) | ~ p105(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ sP8(sK42(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815 | spl63_11307)), inference(subsumption_resolution, [], [f373119, f67765])).
fof(f373119, plain, (p7(sK45(sK42(sK40(sK38(sK36(sK35(sK62))))))) | p106(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ p105(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ sP8(sK42(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815)), inference(resolution, [], [f373070, f279])).
fof(f279, plain, ! [X0] : (r1(X0, sK45(X0)) | p106(X0) | ~ p105(X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f110])).
fof(f373070, plain, (! [X11] : (~ r1(sK42(sK40(sK38(sK36(sK35(sK62))))), X11) | p7(X11)) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815)), inference(subsumption_resolution, [], [f373069, f35503])).
fof(f35503, plain, (sP9(sK40(sK38(sK36(sK35(sK62))))) | ~ spl63_5814), inference(avatar_component_clause, [], [f35502])).
fof(f35502, plain, (spl63_5814 <=> sP9(sK40(sK38(sK36(sK35(sK62)))))), introduced(avatar_definition, [new_symbols(naming, [spl63_5814])])).
fof(f373069, plain, (! [X11] : (~ r1(sK42(sK40(sK38(sK36(sK35(sK62))))), X11) | p7(X11) | ~ sP9(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | spl63_5815)), inference(subsumption_resolution, [], [f373068, f35492])).
fof(f35492, plain, (p104(sK40(sK38(sK36(sK35(sK62))))) | ~ spl63_5811), inference(avatar_component_clause, [], [f35490])).
fof(f35490, plain, (spl63_5811 <=> p104(sK40(sK38(sK36(sK35(sK62)))))), introduced(avatar_definition, [new_symbols(naming, [spl63_5811])])).
fof(f373068, plain, (! [X11] : (~ r1(sK42(sK40(sK38(sK36(sK35(sK62))))), X11) | p7(X11) | ~ p104(sK40(sK38(sK36(sK35(sK62))))) | ~ sP9(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | spl63_5815)), inference(subsumption_resolution, [], [f373048, f35507])).
fof(f35507, plain, (~ p105(sK40(sK38(sK36(sK35(sK62))))) | spl63_5815), inference(avatar_component_clause, [], [f35506])).
fof(f35506, plain, (spl63_5815 <=> p105(sK40(sK38(sK36(sK35(sK62)))))), introduced(avatar_definition, [new_symbols(naming, [spl63_5815])])).
fof(f373048, plain, (! [X11] : (~ r1(sK42(sK40(sK38(sK36(sK35(sK62))))), X11) | p7(X11) | p105(sK40(sK38(sK36(sK35(sK62))))) | ~ p104(sK40(sK38(sK36(sK35(sK62))))) | ~ sP9(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123)), inference(resolution, [], [f373002, f275])).
fof(f275, plain, ! [X0] : (r1(X0, sK42(X0)) | p105(X0) | ~ p104(X0) | ~ sP9(X0)), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ! [X0] : (~ p104(X0) | p105(X0) | ((p105(sK42(X0)) & ~ p106(sK42(X0)) & p6(sK42(X0)) & r1(X0, sK42(X0))) & (p105(sK43(X0)) & ~ p106(sK43(X0)) & ~ p6(sK43(X0)) & r1(X0, sK43(X0)))) | ~ sP9(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42, sK43])], [f102, f104, f103])).
fof(f103, plain, ! [X0] : (? [X1] : (p105(X1) & ~ p106(X1) & p6(X1) & r1(X0, X1)) => (p105(sK42(X0)) & ~ p106(sK42(X0)) & p6(sK42(X0)) & r1(X0, sK42(X0)))), introduced(choice_axiom, [])).
fof(f104, plain, ! [X0] : (? [X2] : (p105(X2) & ~ p106(X2) & ~ p6(X2) & r1(X0, X2)) => (p105(sK43(X0)) & ~ p106(sK43(X0)) & ~ p6(sK43(X0)) & r1(X0, sK43(X0)))), introduced(choice_axiom, [])).
fof(f102, plain, ! [X0] : (~ p104(X0) | p105(X0) | (? [X1] : (p105(X1) & ~ p106(X1) & p6(X1) & r1(X0, X1)) & ? [X2] : (p105(X2) & ~ p106(X2) & ~ p6(X2) & r1(X0, X2))) | ~ sP9(X0)), inference(rectify, [], [f101])).
fof(f101, plain, ! [X15] : (~ p104(X15) | p105(X15) | (? [X56] : (p105(X56) & ~ p106(X56) & p6(X56) & r1(X15, X56)) & ? [X57] : (p105(X57) & ~ p106(X57) & ~ p6(X57) & r1(X15, X57))) | ~ sP9(X15)), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ! [X15] : (~ p104(X15) | p105(X15) | (? [X56] : (p105(X56) & ~ p106(X56) & p6(X56) & r1(X15, X56)) & ? [X57] : (p105(X57) & ~ p106(X57) & ~ p6(X57) & r1(X15, X57))) | ~ sP9(X15)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X15] : (sP9(X15) <=> (~ p104(X15) | p105(X15) | (? [X56] : (p105(X56) & ~ p106(X56) & p6(X56) & r1(X15, X56)) & ? [X57] : (p105(X57) & ~ p106(X57) & ~ p6(X57) & r1(X15, X57))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f373002, plain, (! [X19, X18] : (~ r1(sK40(sK38(sK36(sK35(sK62)))), X18) | ~ r1(X18, X19) | p7(X19)) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123)), inference(subsumption_resolution, [], [f373001, f24126])).
fof(f24126, plain, (sP10(sK38(sK36(sK35(sK62)))) | ~ spl63_4122), inference(avatar_component_clause, [], [f24125])).
fof(f24125, plain, (spl63_4122 <=> sP10(sK38(sK36(sK35(sK62))))), introduced(avatar_definition, [new_symbols(naming, [spl63_4122])])).
fof(f373001, plain, (! [X19, X18] : (~ r1(sK40(sK38(sK36(sK35(sK62)))), X18) | ~ r1(X18, X19) | p7(X19) | ~ sP10(sK38(sK36(sK35(sK62))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | spl63_4123)), inference(subsumption_resolution, [], [f373000, f24115])).
fof(f24115, plain, (p103(sK38(sK36(sK35(sK62)))) | ~ spl63_4119), inference(avatar_component_clause, [], [f24113])).
fof(f24113, plain, (spl63_4119 <=> p103(sK38(sK36(sK35(sK62))))), introduced(avatar_definition, [new_symbols(naming, [spl63_4119])])).
fof(f373000, plain, (! [X19, X18] : (~ r1(sK40(sK38(sK36(sK35(sK62)))), X18) | ~ r1(X18, X19) | p7(X19) | ~ p103(sK38(sK36(sK35(sK62)))) | ~ sP10(sK38(sK36(sK35(sK62))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | spl63_4123)), inference(subsumption_resolution, [], [f372978, f24130])).
fof(f24130, plain, (~ p104(sK38(sK36(sK35(sK62)))) | spl63_4123), inference(avatar_component_clause, [], [f24129])).
fof(f24129, plain, (spl63_4123 <=> p104(sK38(sK36(sK35(sK62))))), introduced(avatar_definition, [new_symbols(naming, [spl63_4123])])).
fof(f372978, plain, (! [X19, X18] : (~ r1(sK40(sK38(sK36(sK35(sK62)))), X18) | ~ r1(X18, X19) | p7(X19) | p104(sK38(sK36(sK35(sK62)))) | ~ p103(sK38(sK36(sK35(sK62)))) | ~ sP10(sK38(sK36(sK35(sK62))))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570)), inference(resolution, [], [f372934, f267])).
fof(f267, plain, ! [X0] : (r1(X0, sK40(X0)) | p104(X0) | ~ p103(X0) | ~ sP10(X0)), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ! [X0] : (~ p103(X0) | p104(X0) | ((p104(sK40(X0)) & ~ p105(sK40(X0)) & p5(sK40(X0)) & r1(X0, sK40(X0))) & (p104(sK41(X0)) & ~ p105(sK41(X0)) & ~ p5(sK41(X0)) & r1(X0, sK41(X0)))) | ~ sP10(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK40, sK41])], [f97, f99, f98])).
fof(f98, plain, ! [X0] : (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) => (p104(sK40(X0)) & ~ p105(sK40(X0)) & p5(sK40(X0)) & r1(X0, sK40(X0)))), introduced(choice_axiom, [])).
fof(f99, plain, ! [X0] : (? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2)) => (p104(sK41(X0)) & ~ p105(sK41(X0)) & ~ p5(sK41(X0)) & r1(X0, sK41(X0)))), introduced(choice_axiom, [])).
fof(f97, plain, ! [X0] : (~ p103(X0) | p104(X0) | (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) & ? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2))) | ~ sP10(X0)), inference(rectify, [], [f96])).
fof(f96, plain, ! [X15] : (~ p103(X15) | p104(X15) | (? [X54] : (p104(X54) & ~ p105(X54) & p5(X54) & r1(X15, X54)) & ? [X55] : (p104(X55) & ~ p105(X55) & ~ p5(X55) & r1(X15, X55))) | ~ sP10(X15)), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ! [X15] : (~ p103(X15) | p104(X15) | (? [X54] : (p104(X54) & ~ p105(X54) & p5(X54) & r1(X15, X54)) & ? [X55] : (p104(X55) & ~ p105(X55) & ~ p5(X55) & r1(X15, X55))) | ~ sP10(X15)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X15] : (sP10(X15) <=> (~ p103(X15) | p104(X15) | (? [X54] : (p104(X54) & ~ p105(X54) & p5(X54) & r1(X15, X54)) & ? [X55] : (p104(X55) & ~ p105(X55) & ~ p5(X55) & r1(X15, X55))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f372934, plain, (! [X23, X21, X22] : (~ r1(sK38(sK36(sK35(sK62))), X23) | ~ r1(X23, X21) | ~ r1(X21, X22) | p7(X22)) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570)), inference(subsumption_resolution, [], [f372933, f10156])).
fof(f10156, plain, (sP11(sK36(sK35(sK62))) | ~ spl63_1569), inference(avatar_component_clause, [], [f10155])).
fof(f10155, plain, (spl63_1569 <=> sP11(sK36(sK35(sK62)))), introduced(avatar_definition, [new_symbols(naming, [spl63_1569])])).
fof(f372933, plain, (! [X23, X21, X22] : (~ r1(X21, X22) | ~ r1(X23, X21) | ~ r1(sK38(sK36(sK35(sK62))), X23) | p7(X22) | ~ sP11(sK36(sK35(sK62)))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | spl63_1570)), inference(subsumption_resolution, [], [f372932, f10145])).
fof(f10145, plain, (p102(sK36(sK35(sK62))) | ~ spl63_1566), inference(avatar_component_clause, [], [f10143])).
fof(f10143, plain, (spl63_1566 <=> p102(sK36(sK35(sK62)))), introduced(avatar_definition, [new_symbols(naming, [spl63_1566])])).
fof(f372932, plain, (! [X23, X21, X22] : (~ r1(X21, X22) | ~ r1(X23, X21) | ~ r1(sK38(sK36(sK35(sK62))), X23) | p7(X22) | ~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62)))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593 | spl63_1570)), inference(subsumption_resolution, [], [f372908, f10160])).
fof(f10160, plain, (~ p103(sK36(sK35(sK62))) | spl63_1570), inference(avatar_component_clause, [], [f10159])).
fof(f10159, plain, (spl63_1570 <=> p103(sK36(sK35(sK62)))), introduced(avatar_definition, [new_symbols(naming, [spl63_1570])])).
fof(f372908, plain, (! [X23, X21, X22] : (~ r1(X21, X22) | ~ r1(X23, X21) | ~ r1(sK38(sK36(sK35(sK62))), X23) | p7(X22) | p103(sK36(sK35(sK62))) | ~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62)))) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593)), inference(resolution, [], [f372866, f259])).
fof(f259, plain, ! [X0] : (r1(X0, sK38(X0)) | p103(X0) | ~ p102(X0) | ~ sP11(X0)), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ! [X0] : (~ p102(X0) | p103(X0) | ((p103(sK38(X0)) & ~ p104(sK38(X0)) & p4(sK38(X0)) & r1(X0, sK38(X0))) & (p103(sK39(X0)) & ~ p104(sK39(X0)) & ~ p4(sK39(X0)) & r1(X0, sK39(X0)))) | ~ sP11(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK38, sK39])], [f92, f94, f93])).
fof(f93, plain, ! [X0] : (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) => (p103(sK38(X0)) & ~ p104(sK38(X0)) & p4(sK38(X0)) & r1(X0, sK38(X0)))), introduced(choice_axiom, [])).
fof(f94, plain, ! [X0] : (? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2)) => (p103(sK39(X0)) & ~ p104(sK39(X0)) & ~ p4(sK39(X0)) & r1(X0, sK39(X0)))), introduced(choice_axiom, [])).
fof(f92, plain, ! [X0] : (~ p102(X0) | p103(X0) | (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) & ? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2))) | ~ sP11(X0)), inference(rectify, [], [f91])).
fof(f91, plain, ! [X15] : (~ p102(X15) | p103(X15) | (? [X52] : (p103(X52) & ~ p104(X52) & p4(X52) & r1(X15, X52)) & ? [X53] : (p103(X53) & ~ p104(X53) & ~ p4(X53) & r1(X15, X53))) | ~ sP11(X15)), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ! [X15] : (~ p102(X15) | p103(X15) | (? [X52] : (p103(X52) & ~ p104(X52) & p4(X52) & r1(X15, X52)) & ? [X53] : (p103(X53) & ~ p104(X53) & ~ p4(X53) & r1(X15, X53))) | ~ sP11(X15)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X15] : (sP11(X15) <=> (~ p102(X15) | p103(X15) | (? [X52] : (p103(X52) & ~ p104(X52) & p4(X52) & r1(X15, X52)) & ? [X53] : (p103(X53) & ~ p104(X53) & ~ p4(X53) & r1(X15, X53))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f372866, plain, (! [X23, X21, X22, X20] : (~ r1(sK36(sK35(sK62)), X23) | ~ r1(X21, X20) | ~ r1(X22, X21) | ~ r1(X23, X22) | p7(X20)) | (~ spl63_139 | ~ spl63_589 | ~ spl63_592 | spl63_593)), inference(subsumption_resolution, [], [f372865, f5152])).
fof(f5152, plain, (sP12(sK35(sK62)) | ~ spl63_592), inference(avatar_component_clause, [], [f5151])).
fof(f5151, plain, (spl63_592 <=> sP12(sK35(sK62))), introduced(avatar_definition, [new_symbols(naming, [spl63_592])])).
fof(f372865, plain, (! [X23, X21, X22, X20] : (p7(X20) | ~ r1(X21, X20) | ~ r1(X22, X21) | ~ r1(X23, X22) | ~ r1(sK36(sK35(sK62)), X23) | ~ sP12(sK35(sK62))) | (~ spl63_139 | ~ spl63_589 | spl63_593)), inference(subsumption_resolution, [], [f372864, f5141])).
fof(f5141, plain, (p101(sK35(sK62)) | ~ spl63_589), inference(avatar_component_clause, [], [f5139])).
fof(f5139, plain, (spl63_589 <=> p101(sK35(sK62))), introduced(avatar_definition, [new_symbols(naming, [spl63_589])])).
fof(f372864, plain, (! [X23, X21, X22, X20] : (p7(X20) | ~ r1(X21, X20) | ~ r1(X22, X21) | ~ r1(X23, X22) | ~ r1(sK36(sK35(sK62)), X23) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62))) | (~ spl63_139 | spl63_593)), inference(subsumption_resolution, [], [f372838, f5156])).
fof(f5156, plain, (~ p102(sK35(sK62)) | spl63_593), inference(avatar_component_clause, [], [f5155])).
fof(f5155, plain, (spl63_593 <=> p102(sK35(sK62))), introduced(avatar_definition, [new_symbols(naming, [spl63_593])])).
fof(f372838, plain, (! [X23, X21, X22, X20] : (p7(X20) | ~ r1(X21, X20) | ~ r1(X22, X21) | ~ r1(X23, X22) | ~ r1(sK36(sK35(sK62)), X23) | p102(sK35(sK62)) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62))) | ~ spl63_139), inference(resolution, [], [f3066, f251])).
fof(f251, plain, ! [X0] : (r1(X0, sK36(X0)) | p102(X0) | ~ p101(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ! [X0] : (~ p101(X0) | p102(X0) | ((p102(sK36(X0)) & ~ p103(sK36(X0)) & p3(sK36(X0)) & r1(X0, sK36(X0))) & (p102(sK37(X0)) & ~ p103(sK37(X0)) & ~ p3(sK37(X0)) & r1(X0, sK37(X0)))) | ~ sP12(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK36, sK37])], [f87, f89, f88])).
fof(f88, plain, ! [X0] : (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) => (p102(sK36(X0)) & ~ p103(sK36(X0)) & p3(sK36(X0)) & r1(X0, sK36(X0)))), introduced(choice_axiom, [])).
fof(f89, plain, ! [X0] : (? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2)) => (p102(sK37(X0)) & ~ p103(sK37(X0)) & ~ p3(sK37(X0)) & r1(X0, sK37(X0)))), introduced(choice_axiom, [])).
fof(f87, plain, ! [X0] : (~ p101(X0) | p102(X0) | (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) & ? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2))) | ~ sP12(X0)), inference(rectify, [], [f86])).
fof(f86, plain, ! [X15] : (~ p101(X15) | p102(X15) | (? [X50] : (p102(X50) & ~ p103(X50) & p3(X50) & r1(X15, X50)) & ? [X51] : (p102(X51) & ~ p103(X51) & ~ p3(X51) & r1(X15, X51))) | ~ sP12(X15)), inference(nnf_transformation, [], [f21])).
fof(f21, plain, ! [X15] : (~ p101(X15) | p102(X15) | (? [X50] : (p102(X50) & ~ p103(X50) & p3(X50) & r1(X15, X50)) & ? [X51] : (p102(X51) & ~ p103(X51) & ~ p3(X51) & r1(X15, X51))) | ~ sP12(X15)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X15] : (sP12(X15) <=> (~ p101(X15) | p102(X15) | (? [X50] : (p102(X50) & ~ p103(X50) & p3(X50) & r1(X15, X50)) & ? [X51] : (p102(X51) & ~ p103(X51) & ~ p3(X51) & r1(X15, X51))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f3066, plain, (! [X24, X23, X21, X22, X20] : (~ r1(sK35(sK62), X20) | p7(X24) | ~ r1(X23, X24) | ~ r1(X22, X23) | ~ r1(X21, X22) | ~ r1(X20, X21)) | ~ spl63_139), inference(avatar_component_clause, [], [f3065])).
fof(f3065, plain, (spl63_139 <=> ! [X20, X22, X21, X23, X24] : (~ r1(sK35(sK62), X20) | p7(X24) | ~ r1(X23, X24) | ~ r1(X22, X23) | ~ r1(X21, X22) | ~ r1(X20, X21))), introduced(avatar_definition, [new_symbols(naming, [spl63_139])])).
fof(f112907, plain, (~ spl63_5811 | ~ spl63_5814 | spl63_5815 | ~ spl63_11307), inference(avatar_contradiction_clause, [], [f112906])).
fof(f112906, plain, ($false | (~ spl63_5811 | ~ spl63_5814 | spl63_5815 | ~ spl63_11307)), inference(subsumption_resolution, [], [f112905, f35503])).
fof(f112905, plain, (~ sP9(sK40(sK38(sK36(sK35(sK62))))) | (~ spl63_5811 | spl63_5815 | ~ spl63_11307)), inference(subsumption_resolution, [], [f112904, f35492])).
fof(f112904, plain, (~ p104(sK40(sK38(sK36(sK35(sK62))))) | ~ sP9(sK40(sK38(sK36(sK35(sK62))))) | (spl63_5815 | ~ spl63_11307)), inference(subsumption_resolution, [], [f112903, f35507])).
fof(f112903, plain, (p105(sK40(sK38(sK36(sK35(sK62))))) | ~ p104(sK40(sK38(sK36(sK35(sK62))))) | ~ sP9(sK40(sK38(sK36(sK35(sK62))))) | ~ spl63_11307), inference(resolution, [], [f67766, f277])).
fof(f277, plain, ! [X0] : (~ p106(sK42(X0)) | p105(X0) | ~ p104(X0) | ~ sP9(X0)), inference(cnf_transformation, [], [f105])).
fof(f67766, plain, (p106(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ spl63_11307), inference(avatar_component_clause, [], [f67764])).
fof(f112709, plain, (spl63_11306 | ~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815), inference(avatar_split_clause, [], [f112610, f35506, f35502, f35490, f24129, f24125, f24113, f10159, f10155, f10143, f5155, f5151, f5139, f3005, f67760])).
fof(f3005, plain, (spl63_129 <=> ! [X16, X18, X17, X19] : (~ r1(sK35(sK62), X16) | sP31(X19) | ~ r1(X18, X19) | ~ r1(X17, X18) | ~ r1(X16, X17))), introduced(avatar_definition, [new_symbols(naming, [spl63_129])])).
fof(f112610, plain, (sP8(sK42(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815)), inference(resolution, [], [f112560, f164])).
fof(f164, plain, ! [X0] : (~ sP31(X0) | sP8(X0)), inference(cnf_transformation, [], [f43])).
fof(f43, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p107(X0) | p106(X0)) & (~ p108(X0) | p107(X0)) & (~ p109(X0) | p108(X0)) & (~ p110(X0) | p109(X0)) & (~ p111(X0) | p110(X0)) & (~ p112(X0) | p111(X0)) & (~ p113(X0) | p112(X0)) & (~ p114(X0) | p113(X0)) & (~ p115(X0) | p114(X0)) & sP30(X0) & sP29(X0) & sP28(X0) & sP27(X0) & sP26(X0) & sP25(X0) & sP24(X0) & sP23(X0) & sP22(X0) & sP21(X0) & sP20(X0) & sP19(X0) & sP18(X0) & sP17(X0) & sP16(X0) & sP15(X0) & sP13(X0) & sP12(X0) & sP11(X0) & sP10(X0) & sP9(X0) & sP8(X0) & sP7(X0) & sP6(X0) & sP5(X0) & sP4(X0) & sP3(X0) & sP2(X0) & sP1(X0) & sP0(X0) & sP14(X0)) | ~ sP31(X0)), inference(rectify, [], [f42])).
fof(f42, plain, ! [X15] : (((~ p101(X15) | p100(X15)) & (~ p102(X15) | p101(X15)) & (~ p103(X15) | p102(X15)) & (~ p104(X15) | p103(X15)) & (~ p105(X15) | p104(X15)) & (~ p106(X15) | p105(X15)) & (~ p107(X15) | p106(X15)) & (~ p108(X15) | p107(X15)) & (~ p109(X15) | p108(X15)) & (~ p110(X15) | p109(X15)) & (~ p111(X15) | p110(X15)) & (~ p112(X15) | p111(X15)) & (~ p113(X15) | p112(X15)) & (~ p114(X15) | p113(X15)) & (~ p115(X15) | p114(X15)) & sP30(X15) & sP29(X15) & sP28(X15) & sP27(X15) & sP26(X15) & sP25(X15) & sP24(X15) & sP23(X15) & sP22(X15) & sP21(X15) & sP20(X15) & sP19(X15) & sP18(X15) & sP17(X15) & sP16(X15) & sP15(X15) & sP13(X15) & sP12(X15) & sP11(X15) & sP10(X15) & sP9(X15) & sP8(X15) & sP7(X15) & sP6(X15) & sP5(X15) & sP4(X15) & sP3(X15) & sP2(X15) & sP1(X15) & sP0(X15) & sP14(X15)) | ~ sP31(X15)), inference(nnf_transformation, [], [f40])).
fof(f40, plain, ! [X15] : (((~ p101(X15) | p100(X15)) & (~ p102(X15) | p101(X15)) & (~ p103(X15) | p102(X15)) & (~ p104(X15) | p103(X15)) & (~ p105(X15) | p104(X15)) & (~ p106(X15) | p105(X15)) & (~ p107(X15) | p106(X15)) & (~ p108(X15) | p107(X15)) & (~ p109(X15) | p108(X15)) & (~ p110(X15) | p109(X15)) & (~ p111(X15) | p110(X15)) & (~ p112(X15) | p111(X15)) & (~ p113(X15) | p112(X15)) & (~ p114(X15) | p113(X15)) & (~ p115(X15) | p114(X15)) & sP30(X15) & sP29(X15) & sP28(X15) & sP27(X15) & sP26(X15) & sP25(X15) & sP24(X15) & sP23(X15) & sP22(X15) & sP21(X15) & sP20(X15) & sP19(X15) & sP18(X15) & sP17(X15) & sP16(X15) & sP15(X15) & sP13(X15) & sP12(X15) & sP11(X15) & sP10(X15) & sP9(X15) & sP8(X15) & sP7(X15) & sP6(X15) & sP5(X15) & sP4(X15) & sP3(X15) & sP2(X15) & sP1(X15) & sP0(X15) & sP14(X15)) | ~ sP31(X15)), inference(usedef, [], [e40])).
fof(e40, plain, ! [X15] : (sP31(X15) <=> ((~ p101(X15) | p100(X15)) & (~ p102(X15) | p101(X15)) & (~ p103(X15) | p102(X15)) & (~ p104(X15) | p103(X15)) & (~ p105(X15) | p104(X15)) & (~ p106(X15) | p105(X15)) & (~ p107(X15) | p106(X15)) & (~ p108(X15) | p107(X15)) & (~ p109(X15) | p108(X15)) & (~ p110(X15) | p109(X15)) & (~ p111(X15) | p110(X15)) & (~ p112(X15) | p111(X15)) & (~ p113(X15) | p112(X15)) & (~ p114(X15) | p113(X15)) & (~ p115(X15) | p114(X15)) & sP30(X15) & sP29(X15) & sP28(X15) & sP27(X15) & sP26(X15) & sP25(X15) & sP24(X15) & sP23(X15) & sP22(X15) & sP21(X15) & sP20(X15) & sP19(X15) & sP18(X15) & sP17(X15) & sP16(X15) & sP15(X15) & sP13(X15) & sP12(X15) & sP11(X15) & sP10(X15) & sP9(X15) & sP8(X15) & sP7(X15) & sP6(X15) & sP5(X15) & sP4(X15) & sP3(X15) & sP2(X15) & sP1(X15) & sP0(X15) & sP14(X15))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f112560, plain, (sP31(sK42(sK40(sK38(sK36(sK35(sK62)))))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | ~ spl63_5814 | spl63_5815)), inference(subsumption_resolution, [], [f112559, f35503])).
fof(f112559, plain, (sP31(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ sP9(sK40(sK38(sK36(sK35(sK62))))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5811 | spl63_5815)), inference(subsumption_resolution, [], [f112558, f35492])).
fof(f112558, plain, (sP31(sK42(sK40(sK38(sK36(sK35(sK62)))))) | ~ p104(sK40(sK38(sK36(sK35(sK62))))) | ~ sP9(sK40(sK38(sK36(sK35(sK62))))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123 | spl63_5815)), inference(subsumption_resolution, [], [f112538, f35507])).
fof(f112538, plain, (sP31(sK42(sK40(sK38(sK36(sK35(sK62)))))) | p105(sK40(sK38(sK36(sK35(sK62))))) | ~ p104(sK40(sK38(sK36(sK35(sK62))))) | ~ sP9(sK40(sK38(sK36(sK35(sK62))))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123)), inference(resolution, [], [f112486, f275])).
fof(f112486, plain, (! [X9] : (~ r1(sK40(sK38(sK36(sK35(sK62)))), X9) | sP31(X9)) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123)), inference(subsumption_resolution, [], [f112485, f24126])).
fof(f112485, plain, (! [X9] : (~ r1(sK40(sK38(sK36(sK35(sK62)))), X9) | sP31(X9) | ~ sP10(sK38(sK36(sK35(sK62))))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | spl63_4123)), inference(subsumption_resolution, [], [f112484, f24115])).
fof(f112484, plain, (! [X9] : (~ r1(sK40(sK38(sK36(sK35(sK62)))), X9) | sP31(X9) | ~ p103(sK38(sK36(sK35(sK62)))) | ~ sP10(sK38(sK36(sK35(sK62))))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | spl63_4123)), inference(subsumption_resolution, [], [f112462, f24130])).
fof(f112462, plain, (! [X9] : (~ r1(sK40(sK38(sK36(sK35(sK62)))), X9) | sP31(X9) | p104(sK38(sK36(sK35(sK62)))) | ~ p103(sK38(sK36(sK35(sK62)))) | ~ sP10(sK38(sK36(sK35(sK62))))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570)), inference(resolution, [], [f112412, f267])).
fof(f112412, plain, (! [X14, X15] : (~ r1(sK38(sK36(sK35(sK62))), X14) | ~ r1(X14, X15) | sP31(X15)) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570)), inference(subsumption_resolution, [], [f112411, f10156])).
fof(f112411, plain, (! [X14, X15] : (~ r1(X14, X15) | ~ r1(sK38(sK36(sK35(sK62))), X14) | sP31(X15) | ~ sP11(sK36(sK35(sK62)))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | spl63_1570)), inference(subsumption_resolution, [], [f112410, f10145])).
fof(f112410, plain, (! [X14, X15] : (~ r1(X14, X15) | ~ r1(sK38(sK36(sK35(sK62))), X14) | sP31(X15) | ~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62)))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593 | spl63_1570)), inference(subsumption_resolution, [], [f112386, f10160])).
fof(f112386, plain, (! [X14, X15] : (~ r1(X14, X15) | ~ r1(sK38(sK36(sK35(sK62))), X14) | sP31(X15) | p103(sK36(sK35(sK62))) | ~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62)))) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593)), inference(resolution, [], [f112338, f259])).
fof(f112338, plain, (! [X17, X15, X16] : (~ r1(sK36(sK35(sK62)), X17) | ~ r1(X16, X15) | ~ r1(X17, X16) | sP31(X15)) | (~ spl63_129 | ~ spl63_589 | ~ spl63_592 | spl63_593)), inference(subsumption_resolution, [], [f112337, f5152])).
fof(f112337, plain, (! [X17, X15, X16] : (sP31(X15) | ~ r1(X16, X15) | ~ r1(X17, X16) | ~ r1(sK36(sK35(sK62)), X17) | ~ sP12(sK35(sK62))) | (~ spl63_129 | ~ spl63_589 | spl63_593)), inference(subsumption_resolution, [], [f112336, f5141])).
fof(f112336, plain, (! [X17, X15, X16] : (sP31(X15) | ~ r1(X16, X15) | ~ r1(X17, X16) | ~ r1(sK36(sK35(sK62)), X17) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62))) | (~ spl63_129 | spl63_593)), inference(subsumption_resolution, [], [f112310, f5156])).
fof(f112310, plain, (! [X17, X15, X16] : (sP31(X15) | ~ r1(X16, X15) | ~ r1(X17, X16) | ~ r1(sK36(sK35(sK62)), X17) | p102(sK35(sK62)) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62))) | ~ spl63_129), inference(resolution, [], [f3006, f251])).
fof(f3006, plain, (! [X19, X17, X18, X16] : (~ r1(sK35(sK62), X16) | sP31(X19) | ~ r1(X18, X19) | ~ r1(X17, X18) | ~ r1(X16, X17)) | ~ spl63_129), inference(avatar_component_clause, [], [f3005])).
fof(f72028, plain, (~ spl63_5811 | ~ spl63_5814 | spl63_5815 | spl63_11303), inference(avatar_contradiction_clause, [], [f72027])).
fof(f72027, plain, ($false | (~ spl63_5811 | ~ spl63_5814 | spl63_5815 | spl63_11303)), inference(subsumption_resolution, [], [f72026, f35503])).
fof(f72026, plain, (~ sP9(sK40(sK38(sK36(sK35(sK62))))) | (~ spl63_5811 | spl63_5815 | spl63_11303)), inference(subsumption_resolution, [], [f72025, f35492])).
fof(f72025, plain, (~ p104(sK40(sK38(sK36(sK35(sK62))))) | ~ sP9(sK40(sK38(sK36(sK35(sK62))))) | (spl63_5815 | spl63_11303)), inference(subsumption_resolution, [], [f72024, f35507])).
fof(f72024, plain, (p105(sK40(sK38(sK36(sK35(sK62))))) | ~ p104(sK40(sK38(sK36(sK35(sK62))))) | ~ sP9(sK40(sK38(sK36(sK35(sK62))))) | spl63_11303), inference(resolution, [], [f67749, f278])).
fof(f278, plain, ! [X0] : (p105(sK42(X0)) | p105(X0) | ~ p104(X0) | ~ sP9(X0)), inference(cnf_transformation, [], [f105])).
fof(f67749, plain, (~ p105(sK42(sK40(sK38(sK36(sK35(sK62)))))) | spl63_11303), inference(avatar_component_clause, [], [f67748])).
fof(f62229, plain, (~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5815), inference(avatar_contradiction_clause, [], [f62228])).
fof(f62228, plain, ($false | (~ spl63_4119 | ~ spl63_4122 | spl63_4123 | ~ spl63_5815)), inference(subsumption_resolution, [], [f62227, f24126])).
fof(f62227, plain, (~ sP10(sK38(sK36(sK35(sK62)))) | (~ spl63_4119 | spl63_4123 | ~ spl63_5815)), inference(subsumption_resolution, [], [f62226, f24115])).
fof(f62226, plain, (~ p103(sK38(sK36(sK35(sK62)))) | ~ sP10(sK38(sK36(sK35(sK62)))) | (spl63_4123 | ~ spl63_5815)), inference(subsumption_resolution, [], [f62225, f24130])).
fof(f62225, plain, (p104(sK38(sK36(sK35(sK62)))) | ~ p103(sK38(sK36(sK35(sK62)))) | ~ sP10(sK38(sK36(sK35(sK62)))) | ~ spl63_5815), inference(resolution, [], [f35508, f269])).
fof(f269, plain, ! [X0] : (~ p105(sK40(X0)) | p104(X0) | ~ p103(X0) | ~ sP10(X0)), inference(cnf_transformation, [], [f100])).
fof(f35508, plain, (p105(sK40(sK38(sK36(sK35(sK62))))) | ~ spl63_5815), inference(avatar_component_clause, [], [f35506])).
fof(f62018, plain, (spl63_5814 | ~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123), inference(avatar_split_clause, [], [f61921, f24129, f24125, f24113, f10159, f10155, f10143, f5155, f5151, f5139, f2981, f35502])).
fof(f2981, plain, (spl63_125 <=> ! [X13, X12, X14] : (~ r1(sK35(sK62), X12) | sP31(X14) | ~ r1(X13, X14) | ~ r1(X12, X13))), introduced(avatar_definition, [new_symbols(naming, [spl63_125])])).
fof(f61921, plain, (sP9(sK40(sK38(sK36(sK35(sK62))))) | (~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123)), inference(resolution, [], [f61876, f165])).
fof(f165, plain, ! [X0] : (~ sP31(X0) | sP9(X0)), inference(cnf_transformation, [], [f43])).
fof(f61876, plain, (sP31(sK40(sK38(sK36(sK35(sK62))))) | (~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | ~ spl63_4122 | spl63_4123)), inference(subsumption_resolution, [], [f61875, f24126])).
fof(f61875, plain, (sP31(sK40(sK38(sK36(sK35(sK62))))) | ~ sP10(sK38(sK36(sK35(sK62)))) | (~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4119 | spl63_4123)), inference(subsumption_resolution, [], [f61874, f24115])).
fof(f61874, plain, (sP31(sK40(sK38(sK36(sK35(sK62))))) | ~ p103(sK38(sK36(sK35(sK62)))) | ~ sP10(sK38(sK36(sK35(sK62)))) | (~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570 | spl63_4123)), inference(subsumption_resolution, [], [f61852, f24130])).
fof(f61852, plain, (sP31(sK40(sK38(sK36(sK35(sK62))))) | p104(sK38(sK36(sK35(sK62)))) | ~ p103(sK38(sK36(sK35(sK62)))) | ~ sP10(sK38(sK36(sK35(sK62)))) | (~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570)), inference(resolution, [], [f61808, f267])).
fof(f61808, plain, (! [X7] : (~ r1(sK38(sK36(sK35(sK62))), X7) | sP31(X7)) | (~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570)), inference(subsumption_resolution, [], [f61807, f10156])).
fof(f61807, plain, (! [X7] : (~ r1(sK38(sK36(sK35(sK62))), X7) | sP31(X7) | ~ sP11(sK36(sK35(sK62)))) | (~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | spl63_1570)), inference(subsumption_resolution, [], [f61806, f10145])).
fof(f61806, plain, (! [X7] : (~ r1(sK38(sK36(sK35(sK62))), X7) | sP31(X7) | ~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62)))) | (~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593 | spl63_1570)), inference(subsumption_resolution, [], [f61782, f10160])).
fof(f61782, plain, (! [X7] : (~ r1(sK38(sK36(sK35(sK62))), X7) | sP31(X7) | p103(sK36(sK35(sK62))) | ~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62)))) | (~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593)), inference(resolution, [], [f61740, f259])).
fof(f61740, plain, (! [X10, X11] : (~ r1(sK36(sK35(sK62)), X11) | ~ r1(X11, X10) | sP31(X10)) | (~ spl63_125 | ~ spl63_589 | ~ spl63_592 | spl63_593)), inference(subsumption_resolution, [], [f61739, f5152])).
fof(f61739, plain, (! [X10, X11] : (sP31(X10) | ~ r1(X11, X10) | ~ r1(sK36(sK35(sK62)), X11) | ~ sP12(sK35(sK62))) | (~ spl63_125 | ~ spl63_589 | spl63_593)), inference(subsumption_resolution, [], [f61738, f5141])).
fof(f61738, plain, (! [X10, X11] : (sP31(X10) | ~ r1(X11, X10) | ~ r1(sK36(sK35(sK62)), X11) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62))) | (~ spl63_125 | spl63_593)), inference(subsumption_resolution, [], [f61712, f5156])).
fof(f61712, plain, (! [X10, X11] : (sP31(X10) | ~ r1(X11, X10) | ~ r1(sK36(sK35(sK62)), X11) | p102(sK35(sK62)) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62))) | ~ spl63_125), inference(resolution, [], [f2982, f251])).
fof(f2982, plain, (! [X14, X12, X13] : (~ r1(sK35(sK62), X12) | sP31(X14) | ~ r1(X13, X14) | ~ r1(X12, X13)) | ~ spl63_125), inference(avatar_component_clause, [], [f2981])).
fof(f37861, plain, (~ spl63_4119 | ~ spl63_4122 | spl63_4123 | spl63_5811), inference(avatar_contradiction_clause, [], [f37860])).
fof(f37860, plain, ($false | (~ spl63_4119 | ~ spl63_4122 | spl63_4123 | spl63_5811)), inference(subsumption_resolution, [], [f37859, f24126])).
fof(f37859, plain, (~ sP10(sK38(sK36(sK35(sK62)))) | (~ spl63_4119 | spl63_4123 | spl63_5811)), inference(subsumption_resolution, [], [f37858, f24115])).
fof(f37858, plain, (~ p103(sK38(sK36(sK35(sK62)))) | ~ sP10(sK38(sK36(sK35(sK62)))) | (spl63_4123 | spl63_5811)), inference(subsumption_resolution, [], [f37857, f24130])).
fof(f37857, plain, (p104(sK38(sK36(sK35(sK62)))) | ~ p103(sK38(sK36(sK35(sK62)))) | ~ sP10(sK38(sK36(sK35(sK62)))) | spl63_5811), inference(resolution, [], [f35491, f270])).
fof(f270, plain, ! [X0] : (p104(sK40(X0)) | p104(X0) | ~ p103(X0) | ~ sP10(X0)), inference(cnf_transformation, [], [f100])).
fof(f35491, plain, (~ p104(sK40(sK38(sK36(sK35(sK62))))) | spl63_5811), inference(avatar_component_clause, [], [f35490])).
fof(f31730, plain, (~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4123), inference(avatar_contradiction_clause, [], [f31729])).
fof(f31729, plain, ($false | (~ spl63_1566 | ~ spl63_1569 | spl63_1570 | ~ spl63_4123)), inference(subsumption_resolution, [], [f31728, f10156])).
fof(f31728, plain, (~ sP11(sK36(sK35(sK62))) | (~ spl63_1566 | spl63_1570 | ~ spl63_4123)), inference(subsumption_resolution, [], [f31727, f10145])).
fof(f31727, plain, (~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62))) | (spl63_1570 | ~ spl63_4123)), inference(subsumption_resolution, [], [f31726, f10160])).
fof(f31726, plain, (p103(sK36(sK35(sK62))) | ~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62))) | ~ spl63_4123), inference(resolution, [], [f24131, f261])).
fof(f261, plain, ! [X0] : (~ p104(sK38(X0)) | p103(X0) | ~ p102(X0) | ~ sP11(X0)), inference(cnf_transformation, [], [f95])).
fof(f24131, plain, (p104(sK38(sK36(sK35(sK62)))) | ~ spl63_4123), inference(avatar_component_clause, [], [f24129])).
fof(f31537, plain, (spl63_4122 | ~ spl63_121 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570), inference(avatar_split_clause, [], [f31442, f10159, f10155, f10143, f5155, f5151, f5139, f2957, f24125])).
fof(f2957, plain, (spl63_121 <=> ! [X9, X8] : (~ r1(sK35(sK62), X8) | sP31(X9) | ~ r1(X8, X9))), introduced(avatar_definition, [new_symbols(naming, [spl63_121])])).
fof(f31442, plain, (sP10(sK38(sK36(sK35(sK62)))) | (~ spl63_121 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570)), inference(resolution, [], [f31390, f166])).
fof(f166, plain, ! [X0] : (~ sP31(X0) | sP10(X0)), inference(cnf_transformation, [], [f43])).
fof(f31390, plain, (sP31(sK38(sK36(sK35(sK62)))) | (~ spl63_121 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | ~ spl63_1569 | spl63_1570)), inference(subsumption_resolution, [], [f31389, f10156])).
fof(f31389, plain, (sP31(sK38(sK36(sK35(sK62)))) | ~ sP11(sK36(sK35(sK62))) | (~ spl63_121 | ~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1566 | spl63_1570)), inference(subsumption_resolution, [], [f31388, f10145])).
fof(f31388, plain, (sP31(sK38(sK36(sK35(sK62)))) | ~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62))) | (~ spl63_121 | ~ spl63_589 | ~ spl63_592 | spl63_593 | spl63_1570)), inference(subsumption_resolution, [], [f31364, f10160])).
fof(f31364, plain, (sP31(sK38(sK36(sK35(sK62)))) | p103(sK36(sK35(sK62))) | ~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62))) | (~ spl63_121 | ~ spl63_589 | ~ spl63_592 | spl63_593)), inference(resolution, [], [f31316, f259])).
fof(f31316, plain, (! [X5] : (~ r1(sK36(sK35(sK62)), X5) | sP31(X5)) | (~ spl63_121 | ~ spl63_589 | ~ spl63_592 | spl63_593)), inference(subsumption_resolution, [], [f31315, f5152])).
fof(f31315, plain, (! [X5] : (sP31(X5) | ~ r1(sK36(sK35(sK62)), X5) | ~ sP12(sK35(sK62))) | (~ spl63_121 | ~ spl63_589 | spl63_593)), inference(subsumption_resolution, [], [f31314, f5141])).
fof(f31314, plain, (! [X5] : (sP31(X5) | ~ r1(sK36(sK35(sK62)), X5) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62))) | (~ spl63_121 | spl63_593)), inference(subsumption_resolution, [], [f31288, f5156])).
fof(f31288, plain, (! [X5] : (sP31(X5) | ~ r1(sK36(sK35(sK62)), X5) | p102(sK35(sK62)) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62))) | ~ spl63_121), inference(resolution, [], [f2958, f251])).
fof(f2958, plain, (! [X8, X9] : (~ r1(sK35(sK62), X8) | sP31(X9) | ~ r1(X8, X9)) | ~ spl63_121), inference(avatar_component_clause, [], [f2957])).
fof(f30923, plain, (~ spl63_1566 | ~ spl63_1569 | spl63_1570 | spl63_4119), inference(avatar_contradiction_clause, [], [f30922])).
fof(f30922, plain, ($false | (~ spl63_1566 | ~ spl63_1569 | spl63_1570 | spl63_4119)), inference(subsumption_resolution, [], [f30921, f10156])).
fof(f30921, plain, (~ sP11(sK36(sK35(sK62))) | (~ spl63_1566 | spl63_1570 | spl63_4119)), inference(subsumption_resolution, [], [f30920, f10145])).
fof(f30920, plain, (~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62))) | (spl63_1570 | spl63_4119)), inference(subsumption_resolution, [], [f30919, f10160])).
fof(f30919, plain, (p103(sK36(sK35(sK62))) | ~ p102(sK36(sK35(sK62))) | ~ sP11(sK36(sK35(sK62))) | spl63_4119), inference(resolution, [], [f24114, f262])).
fof(f262, plain, ! [X0] : (p103(sK38(X0)) | p103(X0) | ~ p102(X0) | ~ sP11(X0)), inference(cnf_transformation, [], [f95])).
fof(f24114, plain, (~ p103(sK38(sK36(sK35(sK62)))) | spl63_4119), inference(avatar_component_clause, [], [f24113])).
fof(f14481, plain, (~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1570), inference(avatar_contradiction_clause, [], [f14480])).
fof(f14480, plain, ($false | (~ spl63_589 | ~ spl63_592 | spl63_593 | ~ spl63_1570)), inference(subsumption_resolution, [], [f14479, f5152])).
fof(f14479, plain, (~ sP12(sK35(sK62)) | (~ spl63_589 | spl63_593 | ~ spl63_1570)), inference(subsumption_resolution, [], [f14478, f5141])).
fof(f14478, plain, (~ p101(sK35(sK62)) | ~ sP12(sK35(sK62)) | (spl63_593 | ~ spl63_1570)), inference(subsumption_resolution, [], [f14477, f5156])).
fof(f14477, plain, (p102(sK35(sK62)) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62)) | ~ spl63_1570), inference(resolution, [], [f10161, f253])).
fof(f253, plain, ! [X0] : (~ p103(sK36(X0)) | p102(X0) | ~ p101(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f90])).
fof(f10161, plain, (p103(sK36(sK35(sK62))) | ~ spl63_1570), inference(avatar_component_clause, [], [f10159])).
fof(f14328, plain, (~ spl63_589 | ~ spl63_592 | spl63_593 | spl63_1566), inference(avatar_contradiction_clause, [], [f14327])).
fof(f14327, plain, ($false | (~ spl63_589 | ~ spl63_592 | spl63_593 | spl63_1566)), inference(subsumption_resolution, [], [f14326, f5152])).
fof(f14326, plain, (~ sP12(sK35(sK62)) | (~ spl63_589 | spl63_593 | spl63_1566)), inference(subsumption_resolution, [], [f14325, f5141])).
fof(f14325, plain, (~ p101(sK35(sK62)) | ~ sP12(sK35(sK62)) | (spl63_593 | spl63_1566)), inference(subsumption_resolution, [], [f14324, f5156])).
fof(f14324, plain, (p102(sK35(sK62)) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62)) | spl63_1566), inference(resolution, [], [f10144, f254])).
fof(f254, plain, ! [X0] : (p102(sK36(X0)) | p102(X0) | ~ p101(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f90])).
fof(f10144, plain, (~ p102(sK36(sK35(sK62))) | spl63_1566), inference(avatar_component_clause, [], [f10143])).
fof(f14255, plain, (spl63_1569 | ~ spl63_117 | ~ spl63_589 | ~ spl63_592 | spl63_593), inference(avatar_split_clause, [], [f14176, f5155, f5151, f5139, f2931, f10155])).
fof(f2931, plain, (spl63_117 <=> ! [X4] : (~ r1(sK35(sK62), X4) | sP31(X4))), introduced(avatar_definition, [new_symbols(naming, [spl63_117])])).
fof(f14176, plain, (sP11(sK36(sK35(sK62))) | (~ spl63_117 | ~ spl63_589 | ~ spl63_592 | spl63_593)), inference(resolution, [], [f14123, f167])).
fof(f167, plain, ! [X0] : (~ sP31(X0) | sP11(X0)), inference(cnf_transformation, [], [f43])).
fof(f14123, plain, (sP31(sK36(sK35(sK62))) | (~ spl63_117 | ~ spl63_589 | ~ spl63_592 | spl63_593)), inference(subsumption_resolution, [], [f14122, f5152])).
fof(f14122, plain, (sP31(sK36(sK35(sK62))) | ~ sP12(sK35(sK62)) | (~ spl63_117 | ~ spl63_589 | spl63_593)), inference(subsumption_resolution, [], [f14121, f5141])).
fof(f14121, plain, (sP31(sK36(sK35(sK62))) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62)) | (~ spl63_117 | spl63_593)), inference(subsumption_resolution, [], [f14095, f5156])).
fof(f14095, plain, (sP31(sK36(sK35(sK62))) | p102(sK35(sK62)) | ~ p101(sK35(sK62)) | ~ sP12(sK35(sK62)) | ~ spl63_117), inference(resolution, [], [f2932, f251])).
fof(f2932, plain, (! [X4] : (~ r1(sK35(sK62), X4) | sP31(X4)) | ~ spl63_117), inference(avatar_component_clause, [], [f2931])).
fof(f5734, plain, (~ spl63_6 | ~ spl63_593), inference(avatar_contradiction_clause, [], [f5733])).
fof(f5733, plain, ($false | (~ spl63_6 | ~ spl63_593)), inference(subsumption_resolution, [], [f5732, f1494])).
fof(f1494, plain, (sP13(sK62) | ~ spl63_6), inference(avatar_component_clause, [], [f1493])).
fof(f1493, plain, (spl63_6 <=> sP13(sK62)), introduced(avatar_definition, [new_symbols(naming, [spl63_6])])).
fof(f5732, plain, (~ sP13(sK62) | ~ spl63_593), inference(subsumption_resolution, [], [f5731, f354])).
fof(f354, plain, p100(sK62), inference(cnf_transformation, [], [f153])).
fof(f153, plain, (p100(sK62) & ~ p101(sK62) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (sP31(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(sK62, X1)) & ! [X16] : (! [X17] : (! [X18] : (! [X19] : (! [X20] : (! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (! [X26] : (! [X27] : (! [X28] : (! [X29] : (! [X30] : (p7(X30) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X26, X27)) | ~ r1(X25, X26)) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(sK62, X16))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK62])], [f151, f152])).
fof(f152, plain, (? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (sP31(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X16] : (! [X17] : (! [X18] : (! [X19] : (! [X20] : (! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (! [X26] : (! [X27] : (! [X28] : (! [X29] : (! [X30] : (p7(X30) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X26, X27)) | ~ r1(X25, X26)) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X0, X16))) => (p100(sK62) & ~ p101(sK62) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (sP31(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(sK62, X1)) & ! [X16] : (! [X17] : (! [X18] : (! [X19] : (! [X20] : (! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (! [X26] : (! [X27] : (! [X28] : (! [X29] : (! [X30] : (p7(X30) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X26, X27)) | ~ r1(X25, X26)) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(sK62, X16)))), introduced(choice_axiom, [])).
fof(f151, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (sP31(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X16] : (! [X17] : (! [X18] : (! [X19] : (! [X20] : (! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (! [X26] : (! [X27] : (! [X28] : (! [X29] : (! [X30] : (p7(X30) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X26, X27)) | ~ r1(X25, X26)) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X0, X16))), inference(rectify, [], [f41])).
fof(f41, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (sP31(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : (! [X83] : (! [X84] : (! [X85] : (! [X86] : (! [X87] : (! [X88] : (! [X89] : (! [X90] : (! [X91] : (! [X92] : (p7(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | ~ r1(X89, X90)) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | ~ r1(X85, X86)) | ~ r1(X84, X85)) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X0, X78))), inference(definition_folding, [], [f8, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9])).
fof(f9, plain, ! [X15] : (~ p113(X15) | p114(X15) | (? [X74] : (p114(X74) & ~ p115(X74) & p15(X74) & r1(X15, X74)) & ? [X75] : (p114(X75) & ~ p115(X75) & ~ p15(X75) & r1(X15, X75))) | ~ sP0(X15)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X15] : (sP0(X15) <=> (~ p113(X15) | p114(X15) | (? [X74] : (p114(X74) & ~ p115(X74) & p15(X74) & r1(X15, X74)) & ? [X75] : (p114(X75) & ~ p115(X75) & ~ p15(X75) & r1(X15, X75))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f10, plain, ! [X15] : (~ p112(X15) | p113(X15) | (? [X72] : (p113(X72) & ~ p114(X72) & p14(X72) & r1(X15, X72)) & ? [X73] : (p113(X73) & ~ p114(X73) & ~ p14(X73) & r1(X15, X73))) | ~ sP1(X15)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X15] : (sP1(X15) <=> (~ p112(X15) | p113(X15) | (? [X72] : (p113(X72) & ~ p114(X72) & p14(X72) & r1(X15, X72)) & ? [X73] : (p113(X73) & ~ p114(X73) & ~ p14(X73) & r1(X15, X73))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f11, plain, ! [X15] : (~ p111(X15) | p112(X15) | (? [X70] : (p112(X70) & ~ p113(X70) & p13(X70) & r1(X15, X70)) & ? [X71] : (p112(X71) & ~ p113(X71) & ~ p13(X71) & r1(X15, X71))) | ~ sP2(X15)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X15] : (sP2(X15) <=> (~ p111(X15) | p112(X15) | (? [X70] : (p112(X70) & ~ p113(X70) & p13(X70) & r1(X15, X70)) & ? [X71] : (p112(X71) & ~ p113(X71) & ~ p13(X71) & r1(X15, X71))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f12, plain, ! [X15] : (~ p110(X15) | p111(X15) | (? [X68] : (p111(X68) & ~ p112(X68) & p12(X68) & r1(X15, X68)) & ? [X69] : (p111(X69) & ~ p112(X69) & ~ p12(X69) & r1(X15, X69))) | ~ sP3(X15)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X15] : (sP3(X15) <=> (~ p110(X15) | p111(X15) | (? [X68] : (p111(X68) & ~ p112(X68) & p12(X68) & r1(X15, X68)) & ? [X69] : (p111(X69) & ~ p112(X69) & ~ p12(X69) & r1(X15, X69))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f13, plain, ! [X15] : (~ p109(X15) | p110(X15) | (? [X66] : (p110(X66) & ~ p111(X66) & p11(X66) & r1(X15, X66)) & ? [X67] : (p110(X67) & ~ p111(X67) & ~ p11(X67) & r1(X15, X67))) | ~ sP4(X15)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X15] : (sP4(X15) <=> (~ p109(X15) | p110(X15) | (? [X66] : (p110(X66) & ~ p111(X66) & p11(X66) & r1(X15, X66)) & ? [X67] : (p110(X67) & ~ p111(X67) & ~ p11(X67) & r1(X15, X67))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f14, plain, ! [X15] : (~ p108(X15) | p109(X15) | (? [X64] : (p109(X64) & ~ p110(X64) & p10(X64) & r1(X15, X64)) & ? [X65] : (p109(X65) & ~ p110(X65) & ~ p10(X65) & r1(X15, X65))) | ~ sP5(X15)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X15] : (sP5(X15) <=> (~ p108(X15) | p109(X15) | (? [X64] : (p109(X64) & ~ p110(X64) & p10(X64) & r1(X15, X64)) & ? [X65] : (p109(X65) & ~ p110(X65) & ~ p10(X65) & r1(X15, X65))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f15, plain, ! [X15] : (~ p107(X15) | p108(X15) | (? [X62] : (p108(X62) & ~ p109(X62) & p9(X62) & r1(X15, X62)) & ? [X63] : (p108(X63) & ~ p109(X63) & ~ p9(X63) & r1(X15, X63))) | ~ sP6(X15)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X15] : (sP6(X15) <=> (~ p107(X15) | p108(X15) | (? [X62] : (p108(X62) & ~ p109(X62) & p9(X62) & r1(X15, X62)) & ? [X63] : (p108(X63) & ~ p109(X63) & ~ p9(X63) & r1(X15, X63))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f16, plain, ! [X15] : (~ p106(X15) | p107(X15) | (? [X60] : (p107(X60) & ~ p108(X60) & p8(X60) & r1(X15, X60)) & ? [X61] : (p107(X61) & ~ p108(X61) & ~ p8(X61) & r1(X15, X61))) | ~ sP7(X15)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X15] : (sP7(X15) <=> (~ p106(X15) | p107(X15) | (? [X60] : (p107(X60) & ~ p108(X60) & p8(X60) & r1(X15, X60)) & ? [X61] : (p107(X61) & ~ p108(X61) & ~ p8(X61) & r1(X15, X61))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f22, plain, ! [X15] : (~ p100(X15) | p101(X15) | (? [X48] : (p101(X48) & ~ p102(X48) & p2(X48) & r1(X15, X48)) & ? [X49] : (p101(X49) & ~ p102(X49) & ~ p2(X49) & r1(X15, X49))) | ~ sP13(X15)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X15] : (sP13(X15) <=> (~ p100(X15) | p101(X15) | (? [X48] : (p101(X48) & ~ p102(X48) & p2(X48) & r1(X15, X48)) & ? [X49] : (p101(X49) & ~ p102(X49) & ~ p2(X49) & r1(X15, X49))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f23, plain, ! [X15] : (~ p114(X15) | p115(X15) | (? [X76] : (p115(X76) & p16(X76) & r1(X15, X76)) & ? [X77] : (p115(X77) & ~ p16(X77) & r1(X15, X77))) | ~ sP14(X15)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X15] : (sP14(X15) <=> (~ p114(X15) | p115(X15) | (? [X76] : (p115(X76) & p16(X76) & r1(X15, X76)) & ? [X77] : (p115(X77) & ~ p16(X77) & r1(X15, X77))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f24, plain, ! [X15] : (~ p115(X15) | ((~ p16(X15) | ! [X46] : (~ p115(X46) | p16(X46) | ~ r1(X15, X46))) & (p16(X15) | ! [X47] : (~ p115(X47) | ~ p16(X47) | ~ r1(X15, X47)))) | ~ sP15(X15)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X15] : (sP15(X15) <=> (~ p115(X15) | ((~ p16(X15) | ! [X46] : (~ p115(X46) | p16(X46) | ~ r1(X15, X46))) & (p16(X15) | ! [X47] : (~ p115(X47) | ~ p16(X47) | ~ r1(X15, X47)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f25, plain, ! [X15] : (~ p114(X15) | ((~ p15(X15) | ! [X44] : (~ p114(X44) | p15(X44) | ~ r1(X15, X44))) & (p15(X15) | ! [X45] : (~ p114(X45) | ~ p15(X45) | ~ r1(X15, X45)))) | ~ sP16(X15)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X15] : (sP16(X15) <=> (~ p114(X15) | ((~ p15(X15) | ! [X44] : (~ p114(X44) | p15(X44) | ~ r1(X15, X44))) & (p15(X15) | ! [X45] : (~ p114(X45) | ~ p15(X45) | ~ r1(X15, X45)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f26, plain, ! [X15] : (~ p113(X15) | ((~ p14(X15) | ! [X42] : (~ p113(X42) | p14(X42) | ~ r1(X15, X42))) & (p14(X15) | ! [X43] : (~ p113(X43) | ~ p14(X43) | ~ r1(X15, X43)))) | ~ sP17(X15)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X15] : (sP17(X15) <=> (~ p113(X15) | ((~ p14(X15) | ! [X42] : (~ p113(X42) | p14(X42) | ~ r1(X15, X42))) & (p14(X15) | ! [X43] : (~ p113(X43) | ~ p14(X43) | ~ r1(X15, X43)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f27, plain, ! [X15] : (~ p112(X15) | ((~ p13(X15) | ! [X40] : (~ p112(X40) | p13(X40) | ~ r1(X15, X40))) & (p13(X15) | ! [X41] : (~ p112(X41) | ~ p13(X41) | ~ r1(X15, X41)))) | ~ sP18(X15)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X15] : (sP18(X15) <=> (~ p112(X15) | ((~ p13(X15) | ! [X40] : (~ p112(X40) | p13(X40) | ~ r1(X15, X40))) & (p13(X15) | ! [X41] : (~ p112(X41) | ~ p13(X41) | ~ r1(X15, X41)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f28, plain, ! [X15] : (~ p111(X15) | ((~ p12(X15) | ! [X38] : (~ p111(X38) | p12(X38) | ~ r1(X15, X38))) & (p12(X15) | ! [X39] : (~ p111(X39) | ~ p12(X39) | ~ r1(X15, X39)))) | ~ sP19(X15)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X15] : (sP19(X15) <=> (~ p111(X15) | ((~ p12(X15) | ! [X38] : (~ p111(X38) | p12(X38) | ~ r1(X15, X38))) & (p12(X15) | ! [X39] : (~ p111(X39) | ~ p12(X39) | ~ r1(X15, X39)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f29, plain, ! [X15] : (~ p110(X15) | ((~ p11(X15) | ! [X36] : (~ p110(X36) | p11(X36) | ~ r1(X15, X36))) & (p11(X15) | ! [X37] : (~ p110(X37) | ~ p11(X37) | ~ r1(X15, X37)))) | ~ sP20(X15)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X15] : (sP20(X15) <=> (~ p110(X15) | ((~ p11(X15) | ! [X36] : (~ p110(X36) | p11(X36) | ~ r1(X15, X36))) & (p11(X15) | ! [X37] : (~ p110(X37) | ~ p11(X37) | ~ r1(X15, X37)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f30, plain, ! [X15] : (~ p109(X15) | ((~ p10(X15) | ! [X34] : (~ p109(X34) | p10(X34) | ~ r1(X15, X34))) & (p10(X15) | ! [X35] : (~ p109(X35) | ~ p10(X35) | ~ r1(X15, X35)))) | ~ sP21(X15)), inference(usedef, [], [e30])).
fof(e30, plain, ! [X15] : (sP21(X15) <=> (~ p109(X15) | ((~ p10(X15) | ! [X34] : (~ p109(X34) | p10(X34) | ~ r1(X15, X34))) & (p10(X15) | ! [X35] : (~ p109(X35) | ~ p10(X35) | ~ r1(X15, X35)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f31, plain, ! [X15] : (~ p108(X15) | ((~ p9(X15) | ! [X32] : (~ p108(X32) | p9(X32) | ~ r1(X15, X32))) & (p9(X15) | ! [X33] : (~ p108(X33) | ~ p9(X33) | ~ r1(X15, X33)))) | ~ sP22(X15)), inference(usedef, [], [e31])).
fof(e31, plain, ! [X15] : (sP22(X15) <=> (~ p108(X15) | ((~ p9(X15) | ! [X32] : (~ p108(X32) | p9(X32) | ~ r1(X15, X32))) & (p9(X15) | ! [X33] : (~ p108(X33) | ~ p9(X33) | ~ r1(X15, X33)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f32, plain, ! [X15] : (~ p107(X15) | ((~ p8(X15) | ! [X30] : (~ p107(X30) | p8(X30) | ~ r1(X15, X30))) & (p8(X15) | ! [X31] : (~ p107(X31) | ~ p8(X31) | ~ r1(X15, X31)))) | ~ sP23(X15)), inference(usedef, [], [e32])).
fof(e32, plain, ! [X15] : (sP23(X15) <=> (~ p107(X15) | ((~ p8(X15) | ! [X30] : (~ p107(X30) | p8(X30) | ~ r1(X15, X30))) & (p8(X15) | ! [X31] : (~ p107(X31) | ~ p8(X31) | ~ r1(X15, X31)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f33, plain, ! [X15] : (~ p106(X15) | ((~ p7(X15) | ! [X28] : (~ p106(X28) | p7(X28) | ~ r1(X15, X28))) & (p7(X15) | ! [X29] : (~ p106(X29) | ~ p7(X29) | ~ r1(X15, X29)))) | ~ sP24(X15)), inference(usedef, [], [e33])).
fof(e33, plain, ! [X15] : (sP24(X15) <=> (~ p106(X15) | ((~ p7(X15) | ! [X28] : (~ p106(X28) | p7(X28) | ~ r1(X15, X28))) & (p7(X15) | ! [X29] : (~ p106(X29) | ~ p7(X29) | ~ r1(X15, X29)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f34, plain, ! [X15] : (~ p105(X15) | ((~ p6(X15) | ! [X26] : (~ p105(X26) | p6(X26) | ~ r1(X15, X26))) & (p6(X15) | ! [X27] : (~ p105(X27) | ~ p6(X27) | ~ r1(X15, X27)))) | ~ sP25(X15)), inference(usedef, [], [e34])).
fof(e34, plain, ! [X15] : (sP25(X15) <=> (~ p105(X15) | ((~ p6(X15) | ! [X26] : (~ p105(X26) | p6(X26) | ~ r1(X15, X26))) & (p6(X15) | ! [X27] : (~ p105(X27) | ~ p6(X27) | ~ r1(X15, X27)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f35, plain, ! [X15] : (~ p104(X15) | ((~ p5(X15) | ! [X24] : (~ p104(X24) | p5(X24) | ~ r1(X15, X24))) & (p5(X15) | ! [X25] : (~ p104(X25) | ~ p5(X25) | ~ r1(X15, X25)))) | ~ sP26(X15)), inference(usedef, [], [e35])).
fof(e35, plain, ! [X15] : (sP26(X15) <=> (~ p104(X15) | ((~ p5(X15) | ! [X24] : (~ p104(X24) | p5(X24) | ~ r1(X15, X24))) & (p5(X15) | ! [X25] : (~ p104(X25) | ~ p5(X25) | ~ r1(X15, X25)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f36, plain, ! [X15] : (~ p103(X15) | ((~ p4(X15) | ! [X22] : (~ p103(X22) | p4(X22) | ~ r1(X15, X22))) & (p4(X15) | ! [X23] : (~ p103(X23) | ~ p4(X23) | ~ r1(X15, X23)))) | ~ sP27(X15)), inference(usedef, [], [e36])).
fof(e36, plain, ! [X15] : (sP27(X15) <=> (~ p103(X15) | ((~ p4(X15) | ! [X22] : (~ p103(X22) | p4(X22) | ~ r1(X15, X22))) & (p4(X15) | ! [X23] : (~ p103(X23) | ~ p4(X23) | ~ r1(X15, X23)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f37, plain, ! [X15] : (~ p102(X15) | ((~ p3(X15) | ! [X20] : (~ p102(X20) | p3(X20) | ~ r1(X15, X20))) & (p3(X15) | ! [X21] : (~ p102(X21) | ~ p3(X21) | ~ r1(X15, X21)))) | ~ sP28(X15)), inference(usedef, [], [e37])).
fof(e37, plain, ! [X15] : (sP28(X15) <=> (~ p102(X15) | ((~ p3(X15) | ! [X20] : (~ p102(X20) | p3(X20) | ~ r1(X15, X20))) & (p3(X15) | ! [X21] : (~ p102(X21) | ~ p3(X21) | ~ r1(X15, X21)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f38, plain, ! [X15] : (~ p101(X15) | ((~ p2(X15) | ! [X18] : (~ p101(X18) | p2(X18) | ~ r1(X15, X18))) & (p2(X15) | ! [X19] : (~ p101(X19) | ~ p2(X19) | ~ r1(X15, X19)))) | ~ sP29(X15)), inference(usedef, [], [e38])).
fof(e38, plain, ! [X15] : (sP29(X15) <=> (~ p101(X15) | ((~ p2(X15) | ! [X18] : (~ p101(X18) | p2(X18) | ~ r1(X15, X18))) & (p2(X15) | ! [X19] : (~ p101(X19) | ~ p2(X19) | ~ r1(X15, X19)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f39, plain, ! [X15] : (~ p100(X15) | ((~ p1(X15) | ! [X16] : (~ p100(X16) | p1(X16) | ~ r1(X15, X16))) & (p1(X15) | ! [X17] : (~ p100(X17) | ~ p1(X17) | ~ r1(X15, X17)))) | ~ sP30(X15)), inference(usedef, [], [e39])).
fof(e39, plain, ! [X15] : (sP30(X15) <=> (~ p100(X15) | ((~ p1(X15) | ! [X16] : (~ p100(X16) | p1(X16) | ~ r1(X15, X16))) & (p1(X15) | ! [X17] : (~ p100(X17) | ~ p1(X17) | ~ r1(X15, X17)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f8, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (((~ p101(X15) | p100(X15)) & (~ p102(X15) | p101(X15)) & (~ p103(X15) | p102(X15)) & (~ p104(X15) | p103(X15)) & (~ p105(X15) | p104(X15)) & (~ p106(X15) | p105(X15)) & (~ p107(X15) | p106(X15)) & (~ p108(X15) | p107(X15)) & (~ p109(X15) | p108(X15)) & (~ p110(X15) | p109(X15)) & (~ p111(X15) | p110(X15)) & (~ p112(X15) | p111(X15)) & (~ p113(X15) | p112(X15)) & (~ p114(X15) | p113(X15)) & (~ p115(X15) | p114(X15)) & (~ p100(X15) | ((~ p1(X15) | ! [X16] : (~ p100(X16) | p1(X16) | ~ r1(X15, X16))) & (p1(X15) | ! [X17] : (~ p100(X17) | ~ p1(X17) | ~ r1(X15, X17))))) & (~ p101(X15) | ((~ p2(X15) | ! [X18] : (~ p101(X18) | p2(X18) | ~ r1(X15, X18))) & (p2(X15) | ! [X19] : (~ p101(X19) | ~ p2(X19) | ~ r1(X15, X19))))) & (~ p102(X15) | ((~ p3(X15) | ! [X20] : (~ p102(X20) | p3(X20) | ~ r1(X15, X20))) & (p3(X15) | ! [X21] : (~ p102(X21) | ~ p3(X21) | ~ r1(X15, X21))))) & (~ p103(X15) | ((~ p4(X15) | ! [X22] : (~ p103(X22) | p4(X22) | ~ r1(X15, X22))) & (p4(X15) | ! [X23] : (~ p103(X23) | ~ p4(X23) | ~ r1(X15, X23))))) & (~ p104(X15) | ((~ p5(X15) | ! [X24] : (~ p104(X24) | p5(X24) | ~ r1(X15, X24))) & (p5(X15) | ! [X25] : (~ p104(X25) | ~ p5(X25) | ~ r1(X15, X25))))) & (~ p105(X15) | ((~ p6(X15) | ! [X26] : (~ p105(X26) | p6(X26) | ~ r1(X15, X26))) & (p6(X15) | ! [X27] : (~ p105(X27) | ~ p6(X27) | ~ r1(X15, X27))))) & (~ p106(X15) | ((~ p7(X15) | ! [X28] : (~ p106(X28) | p7(X28) | ~ r1(X15, X28))) & (p7(X15) | ! [X29] : (~ p106(X29) | ~ p7(X29) | ~ r1(X15, X29))))) & (~ p107(X15) | ((~ p8(X15) | ! [X30] : (~ p107(X30) | p8(X30) | ~ r1(X15, X30))) & (p8(X15) | ! [X31] : (~ p107(X31) | ~ p8(X31) | ~ r1(X15, X31))))) & (~ p108(X15) | ((~ p9(X15) | ! [X32] : (~ p108(X32) | p9(X32) | ~ r1(X15, X32))) & (p9(X15) | ! [X33] : (~ p108(X33) | ~ p9(X33) | ~ r1(X15, X33))))) & (~ p109(X15) | ((~ p10(X15) | ! [X34] : (~ p109(X34) | p10(X34) | ~ r1(X15, X34))) & (p10(X15) | ! [X35] : (~ p109(X35) | ~ p10(X35) | ~ r1(X15, X35))))) & (~ p110(X15) | ((~ p11(X15) | ! [X36] : (~ p110(X36) | p11(X36) | ~ r1(X15, X36))) & (p11(X15) | ! [X37] : (~ p110(X37) | ~ p11(X37) | ~ r1(X15, X37))))) & (~ p111(X15) | ((~ p12(X15) | ! [X38] : (~ p111(X38) | p12(X38) | ~ r1(X15, X38))) & (p12(X15) | ! [X39] : (~ p111(X39) | ~ p12(X39) | ~ r1(X15, X39))))) & (~ p112(X15) | ((~ p13(X15) | ! [X40] : (~ p112(X40) | p13(X40) | ~ r1(X15, X40))) & (p13(X15) | ! [X41] : (~ p112(X41) | ~ p13(X41) | ~ r1(X15, X41))))) & (~ p113(X15) | ((~ p14(X15) | ! [X42] : (~ p113(X42) | p14(X42) | ~ r1(X15, X42))) & (p14(X15) | ! [X43] : (~ p113(X43) | ~ p14(X43) | ~ r1(X15, X43))))) & (~ p114(X15) | ((~ p15(X15) | ! [X44] : (~ p114(X44) | p15(X44) | ~ r1(X15, X44))) & (p15(X15) | ! [X45] : (~ p114(X45) | ~ p15(X45) | ~ r1(X15, X45))))) & (~ p115(X15) | ((~ p16(X15) | ! [X46] : (~ p115(X46) | p16(X46) | ~ r1(X15, X46))) & (p16(X15) | ! [X47] : (~ p115(X47) | ~ p16(X47) | ~ r1(X15, X47))))) & (~ p100(X15) | p101(X15) | (? [X48] : (p101(X48) & ~ p102(X48) & p2(X48) & r1(X15, X48)) & ? [X49] : (p101(X49) & ~ p102(X49) & ~ p2(X49) & r1(X15, X49)))) & (~ p101(X15) | p102(X15) | (? [X50] : (p102(X50) & ~ p103(X50) & p3(X50) & r1(X15, X50)) & ? [X51] : (p102(X51) & ~ p103(X51) & ~ p3(X51) & r1(X15, X51)))) & (~ p102(X15) | p103(X15) | (? [X52] : (p103(X52) & ~ p104(X52) & p4(X52) & r1(X15, X52)) & ? [X53] : (p103(X53) & ~ p104(X53) & ~ p4(X53) & r1(X15, X53)))) & (~ p103(X15) | p104(X15) | (? [X54] : (p104(X54) & ~ p105(X54) & p5(X54) & r1(X15, X54)) & ? [X55] : (p104(X55) & ~ p105(X55) & ~ p5(X55) & r1(X15, X55)))) & (~ p104(X15) | p105(X15) | (? [X56] : (p105(X56) & ~ p106(X56) & p6(X56) & r1(X15, X56)) & ? [X57] : (p105(X57) & ~ p106(X57) & ~ p6(X57) & r1(X15, X57)))) & (~ p105(X15) | p106(X15) | (? [X58] : (p106(X58) & ~ p107(X58) & p7(X58) & r1(X15, X58)) & ? [X59] : (p106(X59) & ~ p107(X59) & ~ p7(X59) & r1(X15, X59)))) & (~ p106(X15) | p107(X15) | (? [X60] : (p107(X60) & ~ p108(X60) & p8(X60) & r1(X15, X60)) & ? [X61] : (p107(X61) & ~ p108(X61) & ~ p8(X61) & r1(X15, X61)))) & (~ p107(X15) | p108(X15) | (? [X62] : (p108(X62) & ~ p109(X62) & p9(X62) & r1(X15, X62)) & ? [X63] : (p108(X63) & ~ p109(X63) & ~ p9(X63) & r1(X15, X63)))) & (~ p108(X15) | p109(X15) | (? [X64] : (p109(X64) & ~ p110(X64) & p10(X64) & r1(X15, X64)) & ? [X65] : (p109(X65) & ~ p110(X65) & ~ p10(X65) & r1(X15, X65)))) & (~ p109(X15) | p110(X15) | (? [X66] : (p110(X66) & ~ p111(X66) & p11(X66) & r1(X15, X66)) & ? [X67] : (p110(X67) & ~ p111(X67) & ~ p11(X67) & r1(X15, X67)))) & (~ p110(X15) | p111(X15) | (? [X68] : (p111(X68) & ~ p112(X68) & p12(X68) & r1(X15, X68)) & ? [X69] : (p111(X69) & ~ p112(X69) & ~ p12(X69) & r1(X15, X69)))) & (~ p111(X15) | p112(X15) | (? [X70] : (p112(X70) & ~ p113(X70) & p13(X70) & r1(X15, X70)) & ? [X71] : (p112(X71) & ~ p113(X71) & ~ p13(X71) & r1(X15, X71)))) & (~ p112(X15) | p113(X15) | (? [X72] : (p113(X72) & ~ p114(X72) & p14(X72) & r1(X15, X72)) & ? [X73] : (p113(X73) & ~ p114(X73) & ~ p14(X73) & r1(X15, X73)))) & (~ p113(X15) | p114(X15) | (? [X74] : (p114(X74) & ~ p115(X74) & p15(X74) & r1(X15, X74)) & ? [X75] : (p114(X75) & ~ p115(X75) & ~ p15(X75) & r1(X15, X75)))) & (~ p114(X15) | p115(X15) | (? [X76] : (p115(X76) & p16(X76) & r1(X15, X76)) & ? [X77] : (p115(X77) & ~ p16(X77) & r1(X15, X77))))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : (! [X83] : (! [X84] : (! [X85] : (! [X86] : (! [X87] : (! [X88] : (! [X89] : (! [X90] : (! [X91] : (! [X92] : (p7(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | ~ r1(X89, X90)) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | ~ r1(X85, X86)) | ~ r1(X84, X85)) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X0, X78))), inference(flattening, [], [f7])).
fof(f7, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (((~ p101(X15) | p100(X15)) & (~ p102(X15) | p101(X15)) & (~ p103(X15) | p102(X15)) & (~ p104(X15) | p103(X15)) & (~ p105(X15) | p104(X15)) & (~ p106(X15) | p105(X15)) & (~ p107(X15) | p106(X15)) & (~ p108(X15) | p107(X15)) & (~ p109(X15) | p108(X15)) & (~ p110(X15) | p109(X15)) & (~ p111(X15) | p110(X15)) & (~ p112(X15) | p111(X15)) & (~ p113(X15) | p112(X15)) & (~ p114(X15) | p113(X15)) & (~ p115(X15) | p114(X15)) & (~ p100(X15) | ((~ p1(X15) | ! [X16] : (~ p100(X16) | p1(X16) | ~ r1(X15, X16))) & (p1(X15) | ! [X17] : (~ p100(X17) | ~ p1(X17) | ~ r1(X15, X17))))) & (~ p101(X15) | ((~ p2(X15) | ! [X18] : (~ p101(X18) | p2(X18) | ~ r1(X15, X18))) & (p2(X15) | ! [X19] : (~ p101(X19) | ~ p2(X19) | ~ r1(X15, X19))))) & (~ p102(X15) | ((~ p3(X15) | ! [X20] : (~ p102(X20) | p3(X20) | ~ r1(X15, X20))) & (p3(X15) | ! [X21] : (~ p102(X21) | ~ p3(X21) | ~ r1(X15, X21))))) & (~ p103(X15) | ((~ p4(X15) | ! [X22] : (~ p103(X22) | p4(X22) | ~ r1(X15, X22))) & (p4(X15) | ! [X23] : (~ p103(X23) | ~ p4(X23) | ~ r1(X15, X23))))) & (~ p104(X15) | ((~ p5(X15) | ! [X24] : (~ p104(X24) | p5(X24) | ~ r1(X15, X24))) & (p5(X15) | ! [X25] : (~ p104(X25) | ~ p5(X25) | ~ r1(X15, X25))))) & (~ p105(X15) | ((~ p6(X15) | ! [X26] : (~ p105(X26) | p6(X26) | ~ r1(X15, X26))) & (p6(X15) | ! [X27] : (~ p105(X27) | ~ p6(X27) | ~ r1(X15, X27))))) & (~ p106(X15) | ((~ p7(X15) | ! [X28] : (~ p106(X28) | p7(X28) | ~ r1(X15, X28))) & (p7(X15) | ! [X29] : (~ p106(X29) | ~ p7(X29) | ~ r1(X15, X29))))) & (~ p107(X15) | ((~ p8(X15) | ! [X30] : (~ p107(X30) | p8(X30) | ~ r1(X15, X30))) & (p8(X15) | ! [X31] : (~ p107(X31) | ~ p8(X31) | ~ r1(X15, X31))))) & (~ p108(X15) | ((~ p9(X15) | ! [X32] : (~ p108(X32) | p9(X32) | ~ r1(X15, X32))) & (p9(X15) | ! [X33] : (~ p108(X33) | ~ p9(X33) | ~ r1(X15, X33))))) & (~ p109(X15) | ((~ p10(X15) | ! [X34] : (~ p109(X34) | p10(X34) | ~ r1(X15, X34))) & (p10(X15) | ! [X35] : (~ p109(X35) | ~ p10(X35) | ~ r1(X15, X35))))) & (~ p110(X15) | ((~ p11(X15) | ! [X36] : (~ p110(X36) | p11(X36) | ~ r1(X15, X36))) & (p11(X15) | ! [X37] : (~ p110(X37) | ~ p11(X37) | ~ r1(X15, X37))))) & (~ p111(X15) | ((~ p12(X15) | ! [X38] : (~ p111(X38) | p12(X38) | ~ r1(X15, X38))) & (p12(X15) | ! [X39] : (~ p111(X39) | ~ p12(X39) | ~ r1(X15, X39))))) & (~ p112(X15) | ((~ p13(X15) | ! [X40] : (~ p112(X40) | p13(X40) | ~ r1(X15, X40))) & (p13(X15) | ! [X41] : (~ p112(X41) | ~ p13(X41) | ~ r1(X15, X41))))) & (~ p113(X15) | ((~ p14(X15) | ! [X42] : (~ p113(X42) | p14(X42) | ~ r1(X15, X42))) & (p14(X15) | ! [X43] : (~ p113(X43) | ~ p14(X43) | ~ r1(X15, X43))))) & (~ p114(X15) | ((~ p15(X15) | ! [X44] : (~ p114(X44) | p15(X44) | ~ r1(X15, X44))) & (p15(X15) | ! [X45] : (~ p114(X45) | ~ p15(X45) | ~ r1(X15, X45))))) & (~ p115(X15) | ((~ p16(X15) | ! [X46] : (~ p115(X46) | p16(X46) | ~ r1(X15, X46))) & (p16(X15) | ! [X47] : (~ p115(X47) | ~ p16(X47) | ~ r1(X15, X47))))) & ((~ p100(X15) | p101(X15)) | (? [X48] : ((p101(X48) & ~ p102(X48) & p2(X48)) & r1(X15, X48)) & ? [X49] : ((p101(X49) & ~ p102(X49) & ~ p2(X49)) & r1(X15, X49)))) & ((~ p101(X15) | p102(X15)) | (? [X50] : ((p102(X50) & ~ p103(X50) & p3(X50)) & r1(X15, X50)) & ? [X51] : ((p102(X51) & ~ p103(X51) & ~ p3(X51)) & r1(X15, X51)))) & ((~ p102(X15) | p103(X15)) | (? [X52] : ((p103(X52) & ~ p104(X52) & p4(X52)) & r1(X15, X52)) & ? [X53] : ((p103(X53) & ~ p104(X53) & ~ p4(X53)) & r1(X15, X53)))) & ((~ p103(X15) | p104(X15)) | (? [X54] : ((p104(X54) & ~ p105(X54) & p5(X54)) & r1(X15, X54)) & ? [X55] : ((p104(X55) & ~ p105(X55) & ~ p5(X55)) & r1(X15, X55)))) & ((~ p104(X15) | p105(X15)) | (? [X56] : ((p105(X56) & ~ p106(X56) & p6(X56)) & r1(X15, X56)) & ? [X57] : ((p105(X57) & ~ p106(X57) & ~ p6(X57)) & r1(X15, X57)))) & ((~ p105(X15) | p106(X15)) | (? [X58] : ((p106(X58) & ~ p107(X58) & p7(X58)) & r1(X15, X58)) & ? [X59] : ((p106(X59) & ~ p107(X59) & ~ p7(X59)) & r1(X15, X59)))) & ((~ p106(X15) | p107(X15)) | (? [X60] : ((p107(X60) & ~ p108(X60) & p8(X60)) & r1(X15, X60)) & ? [X61] : ((p107(X61) & ~ p108(X61) & ~ p8(X61)) & r1(X15, X61)))) & ((~ p107(X15) | p108(X15)) | (? [X62] : ((p108(X62) & ~ p109(X62) & p9(X62)) & r1(X15, X62)) & ? [X63] : ((p108(X63) & ~ p109(X63) & ~ p9(X63)) & r1(X15, X63)))) & ((~ p108(X15) | p109(X15)) | (? [X64] : ((p109(X64) & ~ p110(X64) & p10(X64)) & r1(X15, X64)) & ? [X65] : ((p109(X65) & ~ p110(X65) & ~ p10(X65)) & r1(X15, X65)))) & ((~ p109(X15) | p110(X15)) | (? [X66] : ((p110(X66) & ~ p111(X66) & p11(X66)) & r1(X15, X66)) & ? [X67] : ((p110(X67) & ~ p111(X67) & ~ p11(X67)) & r1(X15, X67)))) & ((~ p110(X15) | p111(X15)) | (? [X68] : ((p111(X68) & ~ p112(X68) & p12(X68)) & r1(X15, X68)) & ? [X69] : ((p111(X69) & ~ p112(X69) & ~ p12(X69)) & r1(X15, X69)))) & ((~ p111(X15) | p112(X15)) | (? [X70] : ((p112(X70) & ~ p113(X70) & p13(X70)) & r1(X15, X70)) & ? [X71] : ((p112(X71) & ~ p113(X71) & ~ p13(X71)) & r1(X15, X71)))) & ((~ p112(X15) | p113(X15)) | (? [X72] : ((p113(X72) & ~ p114(X72) & p14(X72)) & r1(X15, X72)) & ? [X73] : ((p113(X73) & ~ p114(X73) & ~ p14(X73)) & r1(X15, X73)))) & ((~ p113(X15) | p114(X15)) | (? [X74] : ((p114(X74) & ~ p115(X74) & p15(X74)) & r1(X15, X74)) & ? [X75] : ((p114(X75) & ~ p115(X75) & ~ p15(X75)) & r1(X15, X75)))) & ((~ p114(X15) | p115(X15)) | (? [X76] : ((p115(X76) & p16(X76)) & r1(X15, X76)) & ? [X77] : ((p115(X77) & ~ p16(X77)) & r1(X15, X77))))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) & ! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : (! [X83] : (! [X84] : (! [X85] : (! [X86] : (! [X87] : (! [X88] : (! [X89] : (! [X90] : (! [X91] : (! [X92] : (p7(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | ~ r1(X89, X90)) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | ~ r1(X85, X86)) | ~ r1(X84, X85)) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X0, X78))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (((~ p101(X15) | p100(X15)) & (~ p102(X15) | p101(X15)) & (~ p103(X15) | p102(X15)) & (~ p104(X15) | p103(X15)) & (~ p105(X15) | p104(X15)) & (~ p106(X15) | p105(X15)) & (~ p107(X15) | p106(X15)) & (~ p108(X15) | p107(X15)) & (~ p109(X15) | p108(X15)) & (~ p110(X15) | p109(X15)) & (~ p111(X15) | p110(X15)) & (~ p112(X15) | p111(X15)) & (~ p113(X15) | p112(X15)) & (~ p114(X15) | p113(X15)) & (~ p115(X15) | p114(X15)) & (~ p100(X15) | ((~ p1(X15) | ! [X16] : (~ p100(X16) | p1(X16) | ~ r1(X15, X16))) & (p1(X15) | ! [X17] : (~ p100(X17) | ~ p1(X17) | ~ r1(X15, X17))))) & (~ p101(X15) | ((~ p2(X15) | ! [X18] : (~ p101(X18) | p2(X18) | ~ r1(X15, X18))) & (p2(X15) | ! [X19] : (~ p101(X19) | ~ p2(X19) | ~ r1(X15, X19))))) & (~ p102(X15) | ((~ p3(X15) | ! [X20] : (~ p102(X20) | p3(X20) | ~ r1(X15, X20))) & (p3(X15) | ! [X21] : (~ p102(X21) | ~ p3(X21) | ~ r1(X15, X21))))) & (~ p103(X15) | ((~ p4(X15) | ! [X22] : (~ p103(X22) | p4(X22) | ~ r1(X15, X22))) & (p4(X15) | ! [X23] : (~ p103(X23) | ~ p4(X23) | ~ r1(X15, X23))))) & (~ p104(X15) | ((~ p5(X15) | ! [X24] : (~ p104(X24) | p5(X24) | ~ r1(X15, X24))) & (p5(X15) | ! [X25] : (~ p104(X25) | ~ p5(X25) | ~ r1(X15, X25))))) & (~ p105(X15) | ((~ p6(X15) | ! [X26] : (~ p105(X26) | p6(X26) | ~ r1(X15, X26))) & (p6(X15) | ! [X27] : (~ p105(X27) | ~ p6(X27) | ~ r1(X15, X27))))) & (~ p106(X15) | ((~ p7(X15) | ! [X28] : (~ p106(X28) | p7(X28) | ~ r1(X15, X28))) & (p7(X15) | ! [X29] : (~ p106(X29) | ~ p7(X29) | ~ r1(X15, X29))))) & (~ p107(X15) | ((~ p8(X15) | ! [X30] : (~ p107(X30) | p8(X30) | ~ r1(X15, X30))) & (p8(X15) | ! [X31] : (~ p107(X31) | ~ p8(X31) | ~ r1(X15, X31))))) & (~ p108(X15) | ((~ p9(X15) | ! [X32] : (~ p108(X32) | p9(X32) | ~ r1(X15, X32))) & (p9(X15) | ! [X33] : (~ p108(X33) | ~ p9(X33) | ~ r1(X15, X33))))) & (~ p109(X15) | ((~ p10(X15) | ! [X34] : (~ p109(X34) | p10(X34) | ~ r1(X15, X34))) & (p10(X15) | ! [X35] : (~ p109(X35) | ~ p10(X35) | ~ r1(X15, X35))))) & (~ p110(X15) | ((~ p11(X15) | ! [X36] : (~ p110(X36) | p11(X36) | ~ r1(X15, X36))) & (p11(X15) | ! [X37] : (~ p110(X37) | ~ p11(X37) | ~ r1(X15, X37))))) & (~ p111(X15) | ((~ p12(X15) | ! [X38] : (~ p111(X38) | p12(X38) | ~ r1(X15, X38))) & (p12(X15) | ! [X39] : (~ p111(X39) | ~ p12(X39) | ~ r1(X15, X39))))) & (~ p112(X15) | ((~ p13(X15) | ! [X40] : (~ p112(X40) | p13(X40) | ~ r1(X15, X40))) & (p13(X15) | ! [X41] : (~ p112(X41) | ~ p13(X41) | ~ r1(X15, X41))))) & (~ p113(X15) | ((~ p14(X15) | ! [X42] : (~ p113(X42) | p14(X42) | ~ r1(X15, X42))) & (p14(X15) | ! [X43] : (~ p113(X43) | ~ p14(X43) | ~ r1(X15, X43))))) & (~ p114(X15) | ((~ p15(X15) | ! [X44] : (~ p114(X44) | p15(X44) | ~ r1(X15, X44))) & (p15(X15) | ! [X45] : (~ p114(X45) | ~ p15(X45) | ~ r1(X15, X45))))) & (~ p115(X15) | ((~ p16(X15) | ! [X46] : (~ p115(X46) | p16(X46) | ~ r1(X15, X46))) & (p16(X15) | ! [X47] : (~ p115(X47) | ~ p16(X47) | ~ r1(X15, X47))))) & (~ (p100(X15) & ~ p101(X15)) | (~ ! [X48] : (~ (p101(X48) & ~ p102(X48) & p2(X48)) | ~ r1(X15, X48)) & ~ ! [X49] : (~ (p101(X49) & ~ p102(X49) & ~ p2(X49)) | ~ r1(X15, X49)))) & (~ (p101(X15) & ~ p102(X15)) | (~ ! [X50] : (~ (p102(X50) & ~ p103(X50) & p3(X50)) | ~ r1(X15, X50)) & ~ ! [X51] : (~ (p102(X51) & ~ p103(X51) & ~ p3(X51)) | ~ r1(X15, X51)))) & (~ (p102(X15) & ~ p103(X15)) | (~ ! [X52] : (~ (p103(X52) & ~ p104(X52) & p4(X52)) | ~ r1(X15, X52)) & ~ ! [X53] : (~ (p103(X53) & ~ p104(X53) & ~ p4(X53)) | ~ r1(X15, X53)))) & (~ (p103(X15) & ~ p104(X15)) | (~ ! [X54] : (~ (p104(X54) & ~ p105(X54) & p5(X54)) | ~ r1(X15, X54)) & ~ ! [X55] : (~ (p104(X55) & ~ p105(X55) & ~ p5(X55)) | ~ r1(X15, X55)))) & (~ (p104(X15) & ~ p105(X15)) | (~ ! [X56] : (~ (p105(X56) & ~ p106(X56) & p6(X56)) | ~ r1(X15, X56)) & ~ ! [X57] : (~ (p105(X57) & ~ p106(X57) & ~ p6(X57)) | ~ r1(X15, X57)))) & (~ (p105(X15) & ~ p106(X15)) | (~ ! [X58] : (~ (p106(X58) & ~ p107(X58) & p7(X58)) | ~ r1(X15, X58)) & ~ ! [X59] : (~ (p106(X59) & ~ p107(X59) & ~ p7(X59)) | ~ r1(X15, X59)))) & (~ (p106(X15) & ~ p107(X15)) | (~ ! [X60] : (~ (p107(X60) & ~ p108(X60) & p8(X60)) | ~ r1(X15, X60)) & ~ ! [X61] : (~ (p107(X61) & ~ p108(X61) & ~ p8(X61)) | ~ r1(X15, X61)))) & (~ (p107(X15) & ~ p108(X15)) | (~ ! [X62] : (~ (p108(X62) & ~ p109(X62) & p9(X62)) | ~ r1(X15, X62)) & ~ ! [X63] : (~ (p108(X63) & ~ p109(X63) & ~ p9(X63)) | ~ r1(X15, X63)))) & (~ (p108(X15) & ~ p109(X15)) | (~ ! [X64] : (~ (p109(X64) & ~ p110(X64) & p10(X64)) | ~ r1(X15, X64)) & ~ ! [X65] : (~ (p109(X65) & ~ p110(X65) & ~ p10(X65)) | ~ r1(X15, X65)))) & (~ (p109(X15) & ~ p110(X15)) | (~ ! [X66] : (~ (p110(X66) & ~ p111(X66) & p11(X66)) | ~ r1(X15, X66)) & ~ ! [X67] : (~ (p110(X67) & ~ p111(X67) & ~ p11(X67)) | ~ r1(X15, X67)))) & (~ (p110(X15) & ~ p111(X15)) | (~ ! [X68] : (~ (p111(X68) & ~ p112(X68) & p12(X68)) | ~ r1(X15, X68)) & ~ ! [X69] : (~ (p111(X69) & ~ p112(X69) & ~ p12(X69)) | ~ r1(X15, X69)))) & (~ (p111(X15) & ~ p112(X15)) | (~ ! [X70] : (~ (p112(X70) & ~ p113(X70) & p13(X70)) | ~ r1(X15, X70)) & ~ ! [X71] : (~ (p112(X71) & ~ p113(X71) & ~ p13(X71)) | ~ r1(X15, X71)))) & (~ (p112(X15) & ~ p113(X15)) | (~ ! [X72] : (~ (p113(X72) & ~ p114(X72) & p14(X72)) | ~ r1(X15, X72)) & ~ ! [X73] : (~ (p113(X73) & ~ p114(X73) & ~ p14(X73)) | ~ r1(X15, X73)))) & (~ (p113(X15) & ~ p114(X15)) | (~ ! [X74] : (~ (p114(X74) & ~ p115(X74) & p15(X74)) | ~ r1(X15, X74)) & ~ ! [X75] : (~ (p114(X75) & ~ p115(X75) & ~ p15(X75)) | ~ r1(X15, X75)))) & (~ (p114(X15) & ~ p115(X15)) | (~ ! [X76] : (~ (p115(X76) & p16(X76)) | ~ r1(X15, X76)) & ~ ! [X77] : (~ (p115(X77) & ~ p16(X77)) | ~ r1(X15, X77))))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) | ~ ! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : (! [X83] : (! [X84] : (! [X85] : (! [X86] : (! [X87] : (! [X88] : (! [X89] : (! [X90] : (! [X91] : (! [X92] : (p7(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | ~ r1(X89, X90)) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | ~ r1(X85, X86)) | ~ r1(X84, X85)) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X0, X78))), inference(pure_predicate_removal, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (((~ p101(X15) | p100(X15)) & (~ p102(X15) | p101(X15)) & (~ p103(X15) | p102(X15)) & (~ p104(X15) | p103(X15)) & (~ p105(X15) | p104(X15)) & (~ p106(X15) | p105(X15)) & (~ p107(X15) | p106(X15)) & (~ p108(X15) | p107(X15)) & (~ p109(X15) | p108(X15)) & (~ p110(X15) | p109(X15)) & (~ p111(X15) | p110(X15)) & (~ p112(X15) | p111(X15)) & (~ p113(X15) | p112(X15)) & (~ p114(X15) | p113(X15)) & (~ p115(X15) | p114(X15)) & (~ p116(X15) | p115(X15)) & (~ p100(X15) | ((~ p1(X15) | ! [X16] : (~ p100(X16) | p1(X16) | ~ r1(X15, X16))) & (p1(X15) | ! [X17] : (~ p100(X17) | ~ p1(X17) | ~ r1(X15, X17))))) & (~ p101(X15) | ((~ p2(X15) | ! [X18] : (~ p101(X18) | p2(X18) | ~ r1(X15, X18))) & (p2(X15) | ! [X19] : (~ p101(X19) | ~ p2(X19) | ~ r1(X15, X19))))) & (~ p102(X15) | ((~ p3(X15) | ! [X20] : (~ p102(X20) | p3(X20) | ~ r1(X15, X20))) & (p3(X15) | ! [X21] : (~ p102(X21) | ~ p3(X21) | ~ r1(X15, X21))))) & (~ p103(X15) | ((~ p4(X15) | ! [X22] : (~ p103(X22) | p4(X22) | ~ r1(X15, X22))) & (p4(X15) | ! [X23] : (~ p103(X23) | ~ p4(X23) | ~ r1(X15, X23))))) & (~ p104(X15) | ((~ p5(X15) | ! [X24] : (~ p104(X24) | p5(X24) | ~ r1(X15, X24))) & (p5(X15) | ! [X25] : (~ p104(X25) | ~ p5(X25) | ~ r1(X15, X25))))) & (~ p105(X15) | ((~ p6(X15) | ! [X26] : (~ p105(X26) | p6(X26) | ~ r1(X15, X26))) & (p6(X15) | ! [X27] : (~ p105(X27) | ~ p6(X27) | ~ r1(X15, X27))))) & (~ p106(X15) | ((~ p7(X15) | ! [X28] : (~ p106(X28) | p7(X28) | ~ r1(X15, X28))) & (p7(X15) | ! [X29] : (~ p106(X29) | ~ p7(X29) | ~ r1(X15, X29))))) & (~ p107(X15) | ((~ p8(X15) | ! [X30] : (~ p107(X30) | p8(X30) | ~ r1(X15, X30))) & (p8(X15) | ! [X31] : (~ p107(X31) | ~ p8(X31) | ~ r1(X15, X31))))) & (~ p108(X15) | ((~ p9(X15) | ! [X32] : (~ p108(X32) | p9(X32) | ~ r1(X15, X32))) & (p9(X15) | ! [X33] : (~ p108(X33) | ~ p9(X33) | ~ r1(X15, X33))))) & (~ p109(X15) | ((~ p10(X15) | ! [X34] : (~ p109(X34) | p10(X34) | ~ r1(X15, X34))) & (p10(X15) | ! [X35] : (~ p109(X35) | ~ p10(X35) | ~ r1(X15, X35))))) & (~ p110(X15) | ((~ p11(X15) | ! [X36] : (~ p110(X36) | p11(X36) | ~ r1(X15, X36))) & (p11(X15) | ! [X37] : (~ p110(X37) | ~ p11(X37) | ~ r1(X15, X37))))) & (~ p111(X15) | ((~ p12(X15) | ! [X38] : (~ p111(X38) | p12(X38) | ~ r1(X15, X38))) & (p12(X15) | ! [X39] : (~ p111(X39) | ~ p12(X39) | ~ r1(X15, X39))))) & (~ p112(X15) | ((~ p13(X15) | ! [X40] : (~ p112(X40) | p13(X40) | ~ r1(X15, X40))) & (p13(X15) | ! [X41] : (~ p112(X41) | ~ p13(X41) | ~ r1(X15, X41))))) & (~ p113(X15) | ((~ p14(X15) | ! [X42] : (~ p113(X42) | p14(X42) | ~ r1(X15, X42))) & (p14(X15) | ! [X43] : (~ p113(X43) | ~ p14(X43) | ~ r1(X15, X43))))) & (~ p114(X15) | ((~ p15(X15) | ! [X44] : (~ p114(X44) | p15(X44) | ~ r1(X15, X44))) & (p15(X15) | ! [X45] : (~ p114(X45) | ~ p15(X45) | ~ r1(X15, X45))))) & (~ p115(X15) | ((~ p16(X15) | ! [X46] : (~ p115(X46) | p16(X46) | ~ r1(X15, X46))) & (p16(X15) | ! [X47] : (~ p115(X47) | ~ p16(X47) | ~ r1(X15, X47))))) & (~ (p100(X15) & ~ p101(X15)) | (~ ! [X48] : (~ (p101(X48) & ~ p102(X48) & p2(X48)) | ~ r1(X15, X48)) & ~ ! [X49] : (~ (p101(X49) & ~ p102(X49) & ~ p2(X49)) | ~ r1(X15, X49)))) & (~ (p101(X15) & ~ p102(X15)) | (~ ! [X50] : (~ (p102(X50) & ~ p103(X50) & p3(X50)) | ~ r1(X15, X50)) & ~ ! [X51] : (~ (p102(X51) & ~ p103(X51) & ~ p3(X51)) | ~ r1(X15, X51)))) & (~ (p102(X15) & ~ p103(X15)) | (~ ! [X52] : (~ (p103(X52) & ~ p104(X52) & p4(X52)) | ~ r1(X15, X52)) & ~ ! [X53] : (~ (p103(X53) & ~ p104(X53) & ~ p4(X53)) | ~ r1(X15, X53)))) & (~ (p103(X15) & ~ p104(X15)) | (~ ! [X54] : (~ (p104(X54) & ~ p105(X54) & p5(X54)) | ~ r1(X15, X54)) & ~ ! [X55] : (~ (p104(X55) & ~ p105(X55) & ~ p5(X55)) | ~ r1(X15, X55)))) & (~ (p104(X15) & ~ p105(X15)) | (~ ! [X56] : (~ (p105(X56) & ~ p106(X56) & p6(X56)) | ~ r1(X15, X56)) & ~ ! [X57] : (~ (p105(X57) & ~ p106(X57) & ~ p6(X57)) | ~ r1(X15, X57)))) & (~ (p105(X15) & ~ p106(X15)) | (~ ! [X58] : (~ (p106(X58) & ~ p107(X58) & p7(X58)) | ~ r1(X15, X58)) & ~ ! [X59] : (~ (p106(X59) & ~ p107(X59) & ~ p7(X59)) | ~ r1(X15, X59)))) & (~ (p106(X15) & ~ p107(X15)) | (~ ! [X60] : (~ (p107(X60) & ~ p108(X60) & p8(X60)) | ~ r1(X15, X60)) & ~ ! [X61] : (~ (p107(X61) & ~ p108(X61) & ~ p8(X61)) | ~ r1(X15, X61)))) & (~ (p107(X15) & ~ p108(X15)) | (~ ! [X62] : (~ (p108(X62) & ~ p109(X62) & p9(X62)) | ~ r1(X15, X62)) & ~ ! [X63] : (~ (p108(X63) & ~ p109(X63) & ~ p9(X63)) | ~ r1(X15, X63)))) & (~ (p108(X15) & ~ p109(X15)) | (~ ! [X64] : (~ (p109(X64) & ~ p110(X64) & p10(X64)) | ~ r1(X15, X64)) & ~ ! [X65] : (~ (p109(X65) & ~ p110(X65) & ~ p10(X65)) | ~ r1(X15, X65)))) & (~ (p109(X15) & ~ p110(X15)) | (~ ! [X66] : (~ (p110(X66) & ~ p111(X66) & p11(X66)) | ~ r1(X15, X66)) & ~ ! [X67] : (~ (p110(X67) & ~ p111(X67) & ~ p11(X67)) | ~ r1(X15, X67)))) & (~ (p110(X15) & ~ p111(X15)) | (~ ! [X68] : (~ (p111(X68) & ~ p112(X68) & p12(X68)) | ~ r1(X15, X68)) & ~ ! [X69] : (~ (p111(X69) & ~ p112(X69) & ~ p12(X69)) | ~ r1(X15, X69)))) & (~ (p111(X15) & ~ p112(X15)) | (~ ! [X70] : (~ (p112(X70) & ~ p113(X70) & p13(X70)) | ~ r1(X15, X70)) & ~ ! [X71] : (~ (p112(X71) & ~ p113(X71) & ~ p13(X71)) | ~ r1(X15, X71)))) & (~ (p112(X15) & ~ p113(X15)) | (~ ! [X72] : (~ (p113(X72) & ~ p114(X72) & p14(X72)) | ~ r1(X15, X72)) & ~ ! [X73] : (~ (p113(X73) & ~ p114(X73) & ~ p14(X73)) | ~ r1(X15, X73)))) & (~ (p113(X15) & ~ p114(X15)) | (~ ! [X74] : (~ (p114(X74) & ~ p115(X74) & p15(X74)) | ~ r1(X15, X74)) & ~ ! [X75] : (~ (p114(X75) & ~ p115(X75) & ~ p15(X75)) | ~ r1(X15, X75)))) & (~ (p114(X15) & ~ p115(X15)) | (~ ! [X76] : (~ (p115(X76) & ~ p116(X76) & p16(X76)) | ~ r1(X15, X76)) & ~ ! [X77] : (~ (p115(X77) & ~ p116(X77) & ~ p16(X77)) | ~ r1(X15, X77))))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) | ~ ! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : (! [X83] : (! [X84] : (! [X85] : (! [X86] : (! [X87] : (! [X88] : (! [X89] : (! [X90] : (! [X91] : (! [X92] : (p7(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | ~ r1(X89, X90)) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | ~ r1(X85, X86)) | ~ r1(X84, X85)) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X0, X78))), inference(flattening, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (((~ p101(X15) | p100(X15)) & (~ p102(X15) | p101(X15)) & (~ p103(X15) | p102(X15)) & (~ p104(X15) | p103(X15)) & (~ p105(X15) | p104(X15)) & (~ p106(X15) | p105(X15)) & (~ p107(X15) | p106(X15)) & (~ p108(X15) | p107(X15)) & (~ p109(X15) | p108(X15)) & (~ p110(X15) | p109(X15)) & (~ p111(X15) | p110(X15)) & (~ p112(X15) | p111(X15)) & (~ p113(X15) | p112(X15)) & (~ p114(X15) | p113(X15)) & (~ p115(X15) | p114(X15)) & (~ p116(X15) | p115(X15)) & (~ p100(X15) | ((~ p1(X15) | ! [X16] : (~ p100(X16) | p1(X16) | ~ r1(X15, X16))) & (p1(X15) | ! [X17] : (~ p100(X17) | ~ p1(X17) | ~ r1(X15, X17))))) & (~ p101(X15) | ((~ p2(X15) | ! [X18] : (~ p101(X18) | p2(X18) | ~ r1(X15, X18))) & (p2(X15) | ! [X19] : (~ p101(X19) | ~ p2(X19) | ~ r1(X15, X19))))) & (~ p102(X15) | ((~ p3(X15) | ! [X20] : (~ p102(X20) | p3(X20) | ~ r1(X15, X20))) & (p3(X15) | ! [X21] : (~ p102(X21) | ~ p3(X21) | ~ r1(X15, X21))))) & (~ p103(X15) | ((~ p4(X15) | ! [X22] : (~ p103(X22) | p4(X22) | ~ r1(X15, X22))) & (p4(X15) | ! [X23] : (~ p103(X23) | ~ p4(X23) | ~ r1(X15, X23))))) & (~ p104(X15) | ((~ p5(X15) | ! [X24] : (~ p104(X24) | p5(X24) | ~ r1(X15, X24))) & (p5(X15) | ! [X25] : (~ p104(X25) | ~ p5(X25) | ~ r1(X15, X25))))) & (~ p105(X15) | ((~ p6(X15) | ! [X26] : (~ p105(X26) | p6(X26) | ~ r1(X15, X26))) & (p6(X15) | ! [X27] : (~ p105(X27) | ~ p6(X27) | ~ r1(X15, X27))))) & (~ p106(X15) | ((~ p7(X15) | ! [X28] : (~ p106(X28) | p7(X28) | ~ r1(X15, X28))) & (p7(X15) | ! [X29] : (~ p106(X29) | ~ p7(X29) | ~ r1(X15, X29))))) & (~ p107(X15) | ((~ p8(X15) | ! [X30] : (~ p107(X30) | p8(X30) | ~ r1(X15, X30))) & (p8(X15) | ! [X31] : (~ p107(X31) | ~ p8(X31) | ~ r1(X15, X31))))) & (~ p108(X15) | ((~ p9(X15) | ! [X32] : (~ p108(X32) | p9(X32) | ~ r1(X15, X32))) & (p9(X15) | ! [X33] : (~ p108(X33) | ~ p9(X33) | ~ r1(X15, X33))))) & (~ p109(X15) | ((~ p10(X15) | ! [X34] : (~ p109(X34) | p10(X34) | ~ r1(X15, X34))) & (p10(X15) | ! [X35] : (~ p109(X35) | ~ p10(X35) | ~ r1(X15, X35))))) & (~ p110(X15) | ((~ p11(X15) | ! [X36] : (~ p110(X36) | p11(X36) | ~ r1(X15, X36))) & (p11(X15) | ! [X37] : (~ p110(X37) | ~ p11(X37) | ~ r1(X15, X37))))) & (~ p111(X15) | ((~ p12(X15) | ! [X38] : (~ p111(X38) | p12(X38) | ~ r1(X15, X38))) & (p12(X15) | ! [X39] : (~ p111(X39) | ~ p12(X39) | ~ r1(X15, X39))))) & (~ p112(X15) | ((~ p13(X15) | ! [X40] : (~ p112(X40) | p13(X40) | ~ r1(X15, X40))) & (p13(X15) | ! [X41] : (~ p112(X41) | ~ p13(X41) | ~ r1(X15, X41))))) & (~ p113(X15) | ((~ p14(X15) | ! [X42] : (~ p113(X42) | p14(X42) | ~ r1(X15, X42))) & (p14(X15) | ! [X43] : (~ p113(X43) | ~ p14(X43) | ~ r1(X15, X43))))) & (~ p114(X15) | ((~ p15(X15) | ! [X44] : (~ p114(X44) | p15(X44) | ~ r1(X15, X44))) & (p15(X15) | ! [X45] : (~ p114(X45) | ~ p15(X45) | ~ r1(X15, X45))))) & (~ p115(X15) | ((~ p16(X15) | ! [X46] : (~ p115(X46) | p16(X46) | ~ r1(X15, X46))) & (p16(X15) | ! [X47] : (~ p115(X47) | ~ p16(X47) | ~ r1(X15, X47))))) & (~ (p100(X15) & ~ p101(X15)) | (~ ! [X48] : (~ (p101(X48) & ~ p102(X48) & p2(X48)) | ~ r1(X15, X48)) & ~ ! [X49] : (~ (p101(X49) & ~ p102(X49) & ~ p2(X49)) | ~ r1(X15, X49)))) & (~ (p101(X15) & ~ p102(X15)) | (~ ! [X50] : (~ (p102(X50) & ~ p103(X50) & p3(X50)) | ~ r1(X15, X50)) & ~ ! [X51] : (~ (p102(X51) & ~ p103(X51) & ~ p3(X51)) | ~ r1(X15, X51)))) & (~ (p102(X15) & ~ p103(X15)) | (~ ! [X52] : (~ (p103(X52) & ~ p104(X52) & p4(X52)) | ~ r1(X15, X52)) & ~ ! [X53] : (~ (p103(X53) & ~ p104(X53) & ~ p4(X53)) | ~ r1(X15, X53)))) & (~ (p103(X15) & ~ p104(X15)) | (~ ! [X54] : (~ (p104(X54) & ~ p105(X54) & p5(X54)) | ~ r1(X15, X54)) & ~ ! [X55] : (~ (p104(X55) & ~ p105(X55) & ~ p5(X55)) | ~ r1(X15, X55)))) & (~ (p104(X15) & ~ p105(X15)) | (~ ! [X56] : (~ (p105(X56) & ~ p106(X56) & p6(X56)) | ~ r1(X15, X56)) & ~ ! [X57] : (~ (p105(X57) & ~ p106(X57) & ~ p6(X57)) | ~ r1(X15, X57)))) & (~ (p105(X15) & ~ p106(X15)) | (~ ! [X58] : (~ (p106(X58) & ~ p107(X58) & p7(X58)) | ~ r1(X15, X58)) & ~ ! [X59] : (~ (p106(X59) & ~ p107(X59) & ~ p7(X59)) | ~ r1(X15, X59)))) & (~ (p106(X15) & ~ p107(X15)) | (~ ! [X60] : (~ (p107(X60) & ~ p108(X60) & p8(X60)) | ~ r1(X15, X60)) & ~ ! [X61] : (~ (p107(X61) & ~ p108(X61) & ~ p8(X61)) | ~ r1(X15, X61)))) & (~ (p107(X15) & ~ p108(X15)) | (~ ! [X62] : (~ (p108(X62) & ~ p109(X62) & p9(X62)) | ~ r1(X15, X62)) & ~ ! [X63] : (~ (p108(X63) & ~ p109(X63) & ~ p9(X63)) | ~ r1(X15, X63)))) & (~ (p108(X15) & ~ p109(X15)) | (~ ! [X64] : (~ (p109(X64) & ~ p110(X64) & p10(X64)) | ~ r1(X15, X64)) & ~ ! [X65] : (~ (p109(X65) & ~ p110(X65) & ~ p10(X65)) | ~ r1(X15, X65)))) & (~ (p109(X15) & ~ p110(X15)) | (~ ! [X66] : (~ (p110(X66) & ~ p111(X66) & p11(X66)) | ~ r1(X15, X66)) & ~ ! [X67] : (~ (p110(X67) & ~ p111(X67) & ~ p11(X67)) | ~ r1(X15, X67)))) & (~ (p110(X15) & ~ p111(X15)) | (~ ! [X68] : (~ (p111(X68) & ~ p112(X68) & p12(X68)) | ~ r1(X15, X68)) & ~ ! [X69] : (~ (p111(X69) & ~ p112(X69) & ~ p12(X69)) | ~ r1(X15, X69)))) & (~ (p111(X15) & ~ p112(X15)) | (~ ! [X70] : (~ (p112(X70) & ~ p113(X70) & p13(X70)) | ~ r1(X15, X70)) & ~ ! [X71] : (~ (p112(X71) & ~ p113(X71) & ~ p13(X71)) | ~ r1(X15, X71)))) & (~ (p112(X15) & ~ p113(X15)) | (~ ! [X72] : (~ (p113(X72) & ~ p114(X72) & p14(X72)) | ~ r1(X15, X72)) & ~ ! [X73] : (~ (p113(X73) & ~ p114(X73) & ~ p14(X73)) | ~ r1(X15, X73)))) & (~ (p113(X15) & ~ p114(X15)) | (~ ! [X74] : (~ (p114(X74) & ~ p115(X74) & p15(X74)) | ~ r1(X15, X74)) & ~ ! [X75] : (~ (p114(X75) & ~ p115(X75) & ~ p15(X75)) | ~ r1(X15, X75)))) & (~ (p114(X15) & ~ p115(X15)) | (~ ! [X76] : (~ (p115(X76) & ~ p116(X76) & p16(X76)) | ~ r1(X15, X76)) & ~ ! [X77] : (~ (p115(X77) & ~ p116(X77) & ~ p16(X77)) | ~ r1(X15, X77))))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) | ~ ! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : (! [X83] : (! [X84] : (! [X85] : (! [X86] : (! [X87] : (! [X88] : (! [X89] : (! [X90] : (! [X91] : (! [X92] : (p7(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | ~ r1(X89, X90)) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | ~ r1(X85, X86)) | ~ r1(X84, X85)) | ~ r1(X83, X84)) | ~ r1(X82, X83)) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X0, X78))), inference(rectify, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ p106(X1) | ((~ p7(X1) | ! [X0] : (~ p106(X0) | p7(X0) | ~ r1(X1, X0))) & (p7(X1) | ! [X0] : (~ p106(X0) | ~ p7(X0) | ~ r1(X1, X0))))) & (~ p107(X1) | ((~ p8(X1) | ! [X0] : (~ p107(X0) | p8(X0) | ~ r1(X1, X0))) & (p8(X1) | ! [X0] : (~ p107(X0) | ~ p8(X0) | ~ r1(X1, X0))))) & (~ p108(X1) | ((~ p9(X1) | ! [X0] : (~ p108(X0) | p9(X0) | ~ r1(X1, X0))) & (p9(X1) | ! [X0] : (~ p108(X0) | ~ p9(X0) | ~ r1(X1, X0))))) & (~ p109(X1) | ((~ p10(X1) | ! [X0] : (~ p109(X0) | p10(X0) | ~ r1(X1, X0))) & (p10(X1) | ! [X0] : (~ p109(X0) | ~ p10(X0) | ~ r1(X1, X0))))) & (~ p110(X1) | ((~ p11(X1) | ! [X0] : (~ p110(X0) | p11(X0) | ~ r1(X1, X0))) & (p11(X1) | ! [X0] : (~ p110(X0) | ~ p11(X0) | ~ r1(X1, X0))))) & (~ p111(X1) | ((~ p12(X1) | ! [X0] : (~ p111(X0) | p12(X0) | ~ r1(X1, X0))) & (p12(X1) | ! [X0] : (~ p111(X0) | ~ p12(X0) | ~ r1(X1, X0))))) & (~ p112(X1) | ((~ p13(X1) | ! [X0] : (~ p112(X0) | p13(X0) | ~ r1(X1, X0))) & (p13(X1) | ! [X0] : (~ p112(X0) | ~ p13(X0) | ~ r1(X1, X0))))) & (~ p113(X1) | ((~ p14(X1) | ! [X0] : (~ p113(X0) | p14(X0) | ~ r1(X1, X0))) & (p14(X1) | ! [X0] : (~ p113(X0) | ~ p14(X0) | ~ r1(X1, X0))))) & (~ p114(X1) | ((~ p15(X1) | ! [X0] : (~ p114(X0) | p15(X0) | ~ r1(X1, X0))) & (p15(X1) | ! [X0] : (~ p114(X0) | ~ p15(X0) | ~ r1(X1, X0))))) & (~ p115(X1) | ((~ p16(X1) | ! [X0] : (~ p115(X0) | p16(X0) | ~ r1(X1, X0))) & (p16(X1) | ! [X0] : (~ p115(X0) | ~ p16(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & p7(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & ~ p7(X0)) | ~ r1(X1, X0)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & p8(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & ~ p8(X0)) | ~ r1(X1, X0)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & p9(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & ~ p9(X0)) | ~ r1(X1, X0)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & p10(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & ~ p10(X0)) | ~ r1(X1, X0)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & p11(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & ~ p11(X0)) | ~ r1(X1, X0)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & p12(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & ~ p12(X0)) | ~ r1(X1, X0)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & p13(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & ~ p13(X0)) | ~ r1(X1, X0)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & p14(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & ~ p14(X0)) | ~ r1(X1, X0)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & p15(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & ~ p15(X0)) | ~ r1(X1, X0)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & p16(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & ~ p16(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (p7(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ p106(X1) | ((~ p7(X1) | ! [X0] : (~ p106(X0) | p7(X0) | ~ r1(X1, X0))) & (p7(X1) | ! [X0] : (~ p106(X0) | ~ p7(X0) | ~ r1(X1, X0))))) & (~ p107(X1) | ((~ p8(X1) | ! [X0] : (~ p107(X0) | p8(X0) | ~ r1(X1, X0))) & (p8(X1) | ! [X0] : (~ p107(X0) | ~ p8(X0) | ~ r1(X1, X0))))) & (~ p108(X1) | ((~ p9(X1) | ! [X0] : (~ p108(X0) | p9(X0) | ~ r1(X1, X0))) & (p9(X1) | ! [X0] : (~ p108(X0) | ~ p9(X0) | ~ r1(X1, X0))))) & (~ p109(X1) | ((~ p10(X1) | ! [X0] : (~ p109(X0) | p10(X0) | ~ r1(X1, X0))) & (p10(X1) | ! [X0] : (~ p109(X0) | ~ p10(X0) | ~ r1(X1, X0))))) & (~ p110(X1) | ((~ p11(X1) | ! [X0] : (~ p110(X0) | p11(X0) | ~ r1(X1, X0))) & (p11(X1) | ! [X0] : (~ p110(X0) | ~ p11(X0) | ~ r1(X1, X0))))) & (~ p111(X1) | ((~ p12(X1) | ! [X0] : (~ p111(X0) | p12(X0) | ~ r1(X1, X0))) & (p12(X1) | ! [X0] : (~ p111(X0) | ~ p12(X0) | ~ r1(X1, X0))))) & (~ p112(X1) | ((~ p13(X1) | ! [X0] : (~ p112(X0) | p13(X0) | ~ r1(X1, X0))) & (p13(X1) | ! [X0] : (~ p112(X0) | ~ p13(X0) | ~ r1(X1, X0))))) & (~ p113(X1) | ((~ p14(X1) | ! [X0] : (~ p113(X0) | p14(X0) | ~ r1(X1, X0))) & (p14(X1) | ! [X0] : (~ p113(X0) | ~ p14(X0) | ~ r1(X1, X0))))) & (~ p114(X1) | ((~ p15(X1) | ! [X0] : (~ p114(X0) | p15(X0) | ~ r1(X1, X0))) & (p15(X1) | ! [X0] : (~ p114(X0) | ~ p15(X0) | ~ r1(X1, X0))))) & (~ p115(X1) | ((~ p16(X1) | ! [X0] : (~ p115(X0) | p16(X0) | ~ r1(X1, X0))) & (p16(X1) | ! [X0] : (~ p115(X0) | ~ p16(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & p7(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & ~ p7(X0)) | ~ r1(X1, X0)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & p8(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & ~ p8(X0)) | ~ r1(X1, X0)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & p9(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & ~ p9(X0)) | ~ r1(X1, X0)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & p10(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & ~ p10(X0)) | ~ r1(X1, X0)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & p11(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & ~ p11(X0)) | ~ r1(X1, X0)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & p12(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & ~ p12(X0)) | ~ r1(X1, X0)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & p13(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & ~ p13(X0)) | ~ r1(X1, X0)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & p14(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & ~ p14(X0)) | ~ r1(X1, X0)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & p15(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & ~ p15(X0)) | ~ r1(X1, X0)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & p16(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & ~ p16(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (p7(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL656+1.015.p', main)).
fof(f5731, plain, (~ p100(sK62) | ~ sP13(sK62) | ~ spl63_593), inference(subsumption_resolution, [], [f5730, f353])).
fof(f353, plain, ~ p101(sK62), inference(cnf_transformation, [], [f153])).
fof(f5730, plain, (p101(sK62) | ~ p100(sK62) | ~ sP13(sK62) | ~ spl63_593), inference(resolution, [], [f5157, f241])).
fof(f241, plain, ! [X0] : (~ p102(sK35(X0)) | p101(X0) | ~ p100(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ! [X0] : (~ p100(X0) | p101(X0) | ((p101(sK34(X0)) & ~ p102(sK34(X0)) & p2(sK34(X0)) & r1(X0, sK34(X0))) & (p101(sK35(X0)) & ~ p102(sK35(X0)) & ~ p2(sK35(X0)) & r1(X0, sK35(X0)))) | ~ sP13(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK34, sK35])], [f82, f84, f83])).
fof(f83, plain, ! [X0] : (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) => (p101(sK34(X0)) & ~ p102(sK34(X0)) & p2(sK34(X0)) & r1(X0, sK34(X0)))), introduced(choice_axiom, [])).
fof(f84, plain, ! [X0] : (? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2)) => (p101(sK35(X0)) & ~ p102(sK35(X0)) & ~ p2(sK35(X0)) & r1(X0, sK35(X0)))), introduced(choice_axiom, [])).
fof(f82, plain, ! [X0] : (~ p100(X0) | p101(X0) | (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) & ? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2))) | ~ sP13(X0)), inference(rectify, [], [f81])).
fof(f81, plain, ! [X15] : (~ p100(X15) | p101(X15) | (? [X48] : (p101(X48) & ~ p102(X48) & p2(X48) & r1(X15, X48)) & ? [X49] : (p101(X49) & ~ p102(X49) & ~ p2(X49) & r1(X15, X49))) | ~ sP13(X15)), inference(nnf_transformation, [], [f22])).
fof(f5157, plain, (p102(sK35(sK62)) | ~ spl63_593), inference(avatar_component_clause, [], [f5155])).
fof(f5537, plain, (spl63_592 | ~ spl63_115), inference(avatar_split_clause, [], [f5455, f2917, f5151])).
fof(f2917, plain, (spl63_115 <=> sP31(sK35(sK62))), introduced(avatar_definition, [new_symbols(naming, [spl63_115])])).
fof(f5455, plain, (sP12(sK35(sK62)) | ~ spl63_115), inference(resolution, [], [f2919, f168])).
fof(f168, plain, ! [X0] : (~ sP31(X0) | sP12(X0)), inference(cnf_transformation, [], [f43])).
fof(f2919, plain, (sP31(sK35(sK62)) | ~ spl63_115), inference(avatar_component_clause, [], [f2917])).
fof(f5433, plain, (~ spl63_6 | spl63_589), inference(avatar_contradiction_clause, [], [f5432])).
fof(f5432, plain, ($false | (~ spl63_6 | spl63_589)), inference(subsumption_resolution, [], [f5431, f1494])).
fof(f5431, plain, (~ sP13(sK62) | spl63_589), inference(subsumption_resolution, [], [f5430, f354])).
fof(f5430, plain, (~ p100(sK62) | ~ sP13(sK62) | spl63_589), inference(subsumption_resolution, [], [f5429, f353])).
fof(f5429, plain, (p101(sK62) | ~ p100(sK62) | ~ sP13(sK62) | spl63_589), inference(resolution, [], [f5140, f242])).
fof(f242, plain, ! [X0] : (p101(sK35(X0)) | p101(X0) | ~ p100(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f85])).
fof(f5140, plain, (~ p101(sK35(sK62)) | spl63_589), inference(avatar_component_clause, [], [f5139])).
fof(f3067, plain, (~ spl63_6 | spl63_139), inference(avatar_split_clause, [], [f3063, f3065, f1493])).
fof(f3063, plain, ! [X24, X23, X21, X22, X20] : (~ r1(sK35(sK62), X20) | ~ r1(X20, X21) | ~ r1(X21, X22) | ~ r1(X22, X23) | ~ r1(X23, X24) | p7(X24) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f3062, f354])).
fof(f3062, plain, ! [X24, X23, X21, X22, X20] : (~ r1(sK35(sK62), X20) | ~ r1(X20, X21) | ~ r1(X21, X22) | ~ r1(X22, X23) | ~ r1(X23, X24) | p7(X24) | ~ p100(sK62) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f2198, f353])).
fof(f2198, plain, ! [X24, X23, X21, X22, X20] : (~ r1(sK35(sK62), X20) | ~ r1(X20, X21) | ~ r1(X21, X22) | ~ r1(X22, X23) | ~ r1(X23, X24) | p7(X24) | p101(sK62) | ~ p100(sK62) | ~ sP13(sK62)), inference(resolution, [], [f2132, f239])).
fof(f239, plain, ! [X0] : (r1(X0, sK35(X0)) | p101(X0) | ~ p100(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f85])).
fof(f2132, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK62, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | p7(X5)), inference(resolution, [], [f2070, f154])).
fof(f154, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL656+1.015.p', reflexivity)).
fof(f2070, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK62, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | p7(X6)), inference(resolution, [], [f2008, f154])).
fof(f2008, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK62, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X7) | p7(X7)), inference(resolution, [], [f1946, f154])).
fof(f1946, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK62, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | ~ r1(X1, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X7) | ~ r1(X7, X8) | p7(X8)), inference(resolution, [], [f1884, f154])).
fof(f1884, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ r1(sK62, X5) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X0, X1) | ~ r1(X1, X6) | ~ r1(X6, X7) | ~ r1(X7, X8) | ~ r1(X8, X9) | p7(X9)), inference(resolution, [], [f1822, f154])).
fof(f1822, plain, ! [X6, X4, X2, X0, X10, X8, X7, X5, X3, X1, X9] : (~ r1(sK62, X7) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X0, X1) | ~ r1(X1, X8) | ~ r1(X8, X9) | ~ r1(X9, X10) | p7(X10)), inference(resolution, [], [f1760, f154])).
fof(f1760, plain, ! [X6, X4, X2, X0, X10, X8, X7, X5, X3, X1, X11, X9] : (~ r1(sK62, X9) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X0, X1) | ~ r1(X1, X10) | ~ r1(X10, X11) | p7(X11)), inference(resolution, [], [f1729, f154])).
fof(f1729, plain, ! [X6, X4, X2, X0, X12, X10, X8, X7, X5, X3, X1, X11, X9] : (~ r1(sK62, X11) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(X11, X10) | ~ r1(X0, X1) | ~ r1(X1, X12) | p7(X12)), inference(resolution, [], [f1439, f154])).
fof(f1439, plain, ! [X6, X4, X2, X0, X12, X10, X8, X7, X5, X3, X1, X13, X11, X9] : (~ r1(sK62, X13) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(X11, X10) | ~ r1(X12, X11) | ~ r1(X13, X12) | ~ r1(X0, X1) | p7(X1)), inference(resolution, [], [f351, f154])).
fof(f351, plain, ! [X30, X28, X26, X24, X23, X21, X19, X17, X29, X27, X25, X22, X20, X18, X16] : (~ r1(sK62, X16) | ~ r1(X29, X30) | ~ r1(X28, X29) | ~ r1(X27, X28) | ~ r1(X26, X27) | ~ r1(X25, X26) | ~ r1(X24, X25) | ~ r1(X23, X24) | ~ r1(X22, X23) | ~ r1(X21, X22) | ~ r1(X20, X21) | ~ r1(X19, X20) | ~ r1(X18, X19) | ~ r1(X17, X18) | ~ r1(X16, X17) | p7(X30)), inference(cnf_transformation, [], [f153])).
fof(f3007, plain, (~ spl63_6 | spl63_129), inference(avatar_split_clause, [], [f3003, f3005, f1493])).
fof(f3003, plain, ! [X19, X17, X18, X16] : (~ r1(sK35(sK62), X16) | ~ r1(X16, X17) | ~ r1(X17, X18) | ~ r1(X18, X19) | sP31(X19) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f3002, f354])).
fof(f3002, plain, ! [X19, X17, X18, X16] : (~ r1(sK35(sK62), X16) | ~ r1(X16, X17) | ~ r1(X17, X18) | ~ r1(X18, X19) | sP31(X19) | ~ p100(sK62) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f2353, f353])).
fof(f2353, plain, ! [X19, X17, X18, X16] : (~ r1(sK35(sK62), X16) | ~ r1(X16, X17) | ~ r1(X17, X18) | ~ r1(X18, X19) | sP31(X19) | p101(sK62) | ~ p100(sK62) | ~ sP13(sK62)), inference(resolution, [], [f2287, f239])).
fof(f2287, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK62, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | sP31(X4)), inference(resolution, [], [f2225, f154])).
fof(f2225, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK62, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | sP31(X5)), inference(resolution, [], [f2163, f154])).
fof(f2163, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK62, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | sP31(X6)), inference(resolution, [], [f2101, f154])).
fof(f2101, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK62, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X7) | sP31(X7)), inference(resolution, [], [f2039, f154])).
fof(f2039, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK62, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | ~ r1(X1, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X7) | ~ r1(X7, X8) | sP31(X8)), inference(resolution, [], [f1977, f154])).
fof(f1977, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ r1(sK62, X5) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X0, X1) | ~ r1(X1, X6) | ~ r1(X6, X7) | ~ r1(X7, X8) | ~ r1(X8, X9) | sP31(X9)), inference(resolution, [], [f1915, f154])).
fof(f1915, plain, ! [X6, X4, X2, X0, X10, X8, X7, X5, X3, X1, X9] : (~ r1(sK62, X7) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X0, X1) | ~ r1(X1, X8) | ~ r1(X8, X9) | ~ r1(X9, X10) | sP31(X10)), inference(resolution, [], [f1853, f154])).
fof(f1853, plain, ! [X6, X4, X2, X0, X10, X8, X7, X5, X3, X1, X11, X9] : (~ r1(sK62, X9) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X0, X1) | ~ r1(X1, X10) | ~ r1(X10, X11) | sP31(X11)), inference(resolution, [], [f1791, f154])).
fof(f1791, plain, ! [X6, X4, X2, X0, X12, X10, X8, X7, X5, X3, X1, X11, X9] : (~ r1(sK62, X11) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(X11, X10) | ~ r1(X0, X1) | ~ r1(X1, X12) | sP31(X12)), inference(resolution, [], [f1698, f154])).
fof(f1698, plain, ! [X6, X4, X2, X0, X12, X10, X8, X7, X5, X3, X1, X13, X11, X9] : (~ r1(sK62, X13) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(X11, X10) | ~ r1(X12, X11) | ~ r1(X13, X12) | ~ r1(X0, X1) | sP31(X1)), inference(resolution, [], [f352, f154])).
fof(f352, plain, ! [X6, X4, X2, X14, X12, X10, X8, X7, X5, X3, X1, X15, X13, X11, X9] : (~ r1(sK62, X1) | ~ r1(X14, X15) | ~ r1(X13, X14) | ~ r1(X12, X13) | ~ r1(X11, X12) | ~ r1(X10, X11) | ~ r1(X9, X10) | ~ r1(X8, X9) | ~ r1(X7, X8) | ~ r1(X6, X7) | ~ r1(X5, X6) | ~ r1(X4, X5) | ~ r1(X3, X4) | ~ r1(X2, X3) | ~ r1(X1, X2) | sP31(X15)), inference(cnf_transformation, [], [f153])).
fof(f2983, plain, (~ spl63_6 | spl63_125), inference(avatar_split_clause, [], [f2979, f2981, f1493])).
fof(f2979, plain, ! [X14, X12, X13] : (~ r1(sK35(sK62), X12) | ~ r1(X12, X13) | ~ r1(X13, X14) | sP31(X14) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f2978, f354])).
fof(f2978, plain, ! [X14, X12, X13] : (~ r1(sK35(sK62), X12) | ~ r1(X12, X13) | ~ r1(X13, X14) | sP31(X14) | ~ p100(sK62) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f2415, f353])).
fof(f2415, plain, ! [X14, X12, X13] : (~ r1(sK35(sK62), X12) | ~ r1(X12, X13) | ~ r1(X13, X14) | sP31(X14) | p101(sK62) | ~ p100(sK62) | ~ sP13(sK62)), inference(resolution, [], [f2349, f239])).
fof(f2349, plain, ! [X2, X0, X3, X1] : (~ r1(sK62, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | sP31(X3)), inference(resolution, [], [f2287, f154])).
fof(f2959, plain, (~ spl63_6 | spl63_121), inference(avatar_split_clause, [], [f2955, f2957, f1493])).
fof(f2955, plain, ! [X8, X9] : (~ r1(sK35(sK62), X8) | ~ r1(X8, X9) | sP31(X9) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f2954, f354])).
fof(f2954, plain, ! [X8, X9] : (~ r1(sK35(sK62), X8) | ~ r1(X8, X9) | sP31(X9) | ~ p100(sK62) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f2477, f353])).
fof(f2477, plain, ! [X8, X9] : (~ r1(sK35(sK62), X8) | ~ r1(X8, X9) | sP31(X9) | p101(sK62) | ~ p100(sK62) | ~ sP13(sK62)), inference(resolution, [], [f2411, f239])).
fof(f2411, plain, ! [X2, X0, X1] : (~ r1(sK62, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP31(X2)), inference(resolution, [], [f2349, f154])).
fof(f2933, plain, (~ spl63_6 | spl63_117), inference(avatar_split_clause, [], [f2929, f2931, f1493])).
fof(f2929, plain, ! [X4] : (~ r1(sK35(sK62), X4) | sP31(X4) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f2928, f354])).
fof(f2928, plain, ! [X4] : (~ r1(sK35(sK62), X4) | sP31(X4) | ~ p100(sK62) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f2539, f353])).
fof(f2539, plain, ! [X4] : (~ r1(sK35(sK62), X4) | sP31(X4) | p101(sK62) | ~ p100(sK62) | ~ sP13(sK62)), inference(resolution, [], [f2473, f239])).
fof(f2473, plain, ! [X0, X1] : (~ r1(sK62, X0) | ~ r1(X0, X1) | sP31(X1)), inference(resolution, [], [f2411, f154])).
fof(f2920, plain, (~ spl63_6 | spl63_115), inference(avatar_split_clause, [], [f2915, f2917, f1493])).
fof(f2915, plain, (sP31(sK35(sK62)) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f2914, f354])).
fof(f2914, plain, (sP31(sK35(sK62)) | ~ p100(sK62) | ~ sP13(sK62)), inference(subsumption_resolution, [], [f2570, f353])).
fof(f2570, plain, (sP31(sK35(sK62)) | p101(sK62) | ~ p100(sK62) | ~ sP13(sK62)), inference(resolution, [], [f2535, f239])).
fof(f2535, plain, ! [X0] : (~ r1(sK62, X0) | sP31(X0)), inference(resolution, [], [f2473, f154])).
fof(f2913, plain, spl63_6, inference(avatar_split_clause, [], [f2611, f1493])).
fof(f2611, plain, sP13(sK62), inference(resolution, [], [f2566, f169])).
fof(f169, plain, ! [X0] : (~ sP31(X0) | sP13(X0)), inference(cnf_transformation, [], [f43])).
fof(f2566, plain, sP31(sK62), inference(resolution, [], [f2535, f154])).