fof(f6760, plain, $false, inference(avatar_sat_refutation, [], [f1669, f2308, f6759])).
fof(f6759, plain, (spl43_29 | ~ spl43_30), inference(avatar_contradiction_clause, [], [f6758])).
fof(f6758, plain, ($false | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f6757, f112])).
fof(f112, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.010.p', reflexivity)).
fof(f6757, plain, (~ r1(sK42, sK42) | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f6756, f247])).
fof(f247, plain, ~ p101(sK42), inference(cnf_transformation, [], [f111])).
fof(f111, plain, (p100(sK42) & ~ p101(sK42) & ! [X1] : (sP21(X1) | ~ r1(sK42, X1)) & ! [X2] : (p5(X2) | ~ r1(sK42, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42])], [f109, f110])).
fof(f110, plain, (? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP21(X1) | ~ r1(X0, X1)) & ! [X2] : (p5(X2) | ~ r1(X0, X2))) => (p100(sK42) & ~ p101(sK42) & ! [X1] : (sP21(X1) | ~ r1(sK42, X1)) & ! [X2] : (p5(X2) | ~ r1(sK42, X2)))), introduced(choice_axiom, [])).
fof(f109, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP21(X1) | ~ r1(X0, X1)) & ! [X2] : (p5(X2) | ~ r1(X0, X2))), inference(rectify, [], [f34])).
fof(f34, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP21(X1) | ~ r1(X0, X1)) & ! [X44] : (p5(X44) | ~ r1(X0, X44))), inference(definition_folding, [], [f11, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12])).
fof(f12, plain, ! [X1] : (~ p108(X1) | p109(X1) | (? [X40] : (p109(X40) & ~ p110(X40) & p10(X40) & r1(X1, X40)) & ? [X41] : (p109(X41) & ~ p110(X41) & ~ p10(X41) & r1(X1, X41))) | ~ sP0(X1)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X1] : (sP0(X1) <=> (~ p108(X1) | p109(X1) | (? [X40] : (p109(X40) & ~ p110(X40) & p10(X40) & r1(X1, X40)) & ? [X41] : (p109(X41) & ~ p110(X41) & ~ p10(X41) & r1(X1, X41))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f13, plain, ! [X1] : (~ p107(X1) | p108(X1) | (? [X38] : (p108(X38) & ~ p109(X38) & p9(X38) & r1(X1, X38)) & ? [X39] : (p108(X39) & ~ p109(X39) & ~ p9(X39) & r1(X1, X39))) | ~ sP1(X1)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X1] : (sP1(X1) <=> (~ p107(X1) | p108(X1) | (? [X38] : (p108(X38) & ~ p109(X38) & p9(X38) & r1(X1, X38)) & ? [X39] : (p108(X39) & ~ p109(X39) & ~ p9(X39) & r1(X1, X39))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f14, plain, ! [X1] : (~ p106(X1) | p107(X1) | (? [X36] : (p107(X36) & ~ p108(X36) & p8(X36) & r1(X1, X36)) & ? [X37] : (p107(X37) & ~ p108(X37) & ~ p8(X37) & r1(X1, X37))) | ~ sP2(X1)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X1] : (sP2(X1) <=> (~ p106(X1) | p107(X1) | (? [X36] : (p107(X36) & ~ p108(X36) & p8(X36) & r1(X1, X36)) & ? [X37] : (p107(X37) & ~ p108(X37) & ~ p8(X37) & r1(X1, X37))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f15, plain, ! [X1] : (~ p105(X1) | p106(X1) | (? [X34] : (p106(X34) & ~ p107(X34) & p7(X34) & r1(X1, X34)) & ? [X35] : (p106(X35) & ~ p107(X35) & ~ p7(X35) & r1(X1, X35))) | ~ sP3(X1)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X1] : (sP3(X1) <=> (~ p105(X1) | p106(X1) | (? [X34] : (p106(X34) & ~ p107(X34) & p7(X34) & r1(X1, X34)) & ? [X35] : (p106(X35) & ~ p107(X35) & ~ p7(X35) & r1(X1, X35))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f16, plain, ! [X1] : (~ p104(X1) | p105(X1) | (? [X32] : (p105(X32) & ~ p106(X32) & p6(X32) & r1(X1, X32)) & ? [X33] : (p105(X33) & ~ p106(X33) & ~ p6(X33) & r1(X1, X33))) | ~ sP4(X1)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X1] : (sP4(X1) <=> (~ p104(X1) | p105(X1) | (? [X32] : (p105(X32) & ~ p106(X32) & p6(X32) & r1(X1, X32)) & ? [X33] : (p105(X33) & ~ p106(X33) & ~ p6(X33) & r1(X1, X33))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f17, plain, ! [X1] : (~ p103(X1) | p104(X1) | (? [X30] : (p104(X30) & ~ p105(X30) & p5(X30) & r1(X1, X30)) & ? [X31] : (p104(X31) & ~ p105(X31) & ~ p5(X31) & r1(X1, X31))) | ~ sP5(X1)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X1] : (sP5(X1) <=> (~ p103(X1) | p104(X1) | (? [X30] : (p104(X30) & ~ p105(X30) & p5(X30) & r1(X1, X30)) & ? [X31] : (p104(X31) & ~ p105(X31) & ~ p5(X31) & r1(X1, X31))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f18, plain, ! [X1] : (~ p102(X1) | p103(X1) | (? [X28] : (p103(X28) & ~ p104(X28) & p4(X28) & r1(X1, X28)) & ? [X29] : (p103(X29) & ~ p104(X29) & ~ p4(X29) & r1(X1, X29))) | ~ sP6(X1)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X1] : (sP6(X1) <=> (~ p102(X1) | p103(X1) | (? [X28] : (p103(X28) & ~ p104(X28) & p4(X28) & r1(X1, X28)) & ? [X29] : (p103(X29) & ~ p104(X29) & ~ p4(X29) & r1(X1, X29))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f19, plain, ! [X1] : (~ p101(X1) | p102(X1) | (? [X26] : (p102(X26) & ~ p103(X26) & p3(X26) & r1(X1, X26)) & ? [X27] : (p102(X27) & ~ p103(X27) & ~ p3(X27) & r1(X1, X27))) | ~ sP7(X1)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X1] : (sP7(X1) <=> (~ p101(X1) | p102(X1) | (? [X26] : (p102(X26) & ~ p103(X26) & p3(X26) & r1(X1, X26)) & ? [X27] : (p102(X27) & ~ p103(X27) & ~ p3(X27) & r1(X1, X27))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f20, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X24] : (p101(X24) & ~ p102(X24) & p2(X24) & r1(X1, X24)) & ? [X25] : (p101(X25) & ~ p102(X25) & ~ p2(X25) & r1(X1, X25))) | ~ sP8(X1)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X1] : (sP8(X1) <=> (~ p100(X1) | p101(X1) | (? [X24] : (p101(X24) & ~ p102(X24) & p2(X24) & r1(X1, X24)) & ? [X25] : (p101(X25) & ~ p102(X25) & ~ p2(X25) & r1(X1, X25))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f21, plain, ! [X1] : (~ p109(X1) | p110(X1) | (? [X42] : (p110(X42) & p11(X42) & r1(X1, X42)) & ? [X43] : (p110(X43) & ~ p11(X43) & r1(X1, X43))) | ~ sP9(X1)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X1] : (sP9(X1) <=> (~ p109(X1) | p110(X1) | (? [X42] : (p110(X42) & p11(X42) & r1(X1, X42)) & ? [X43] : (p110(X43) & ~ p11(X43) & r1(X1, X43))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f22, plain, ! [X1] : (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23)))) | ~ sP10(X1)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X1] : (sP10(X1) <=> (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f23, plain, ! [X1] : (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21)))) | ~ sP11(X1)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X1] : (sP11(X1) <=> (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f24, plain, ! [X1] : (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19)))) | ~ sP12(X1)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X1] : (sP12(X1) <=> (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f25, plain, ! [X1] : (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17)))) | ~ sP13(X1)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X1] : (sP13(X1) <=> (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f26, plain, ! [X1] : (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15)))) | ~ sP14(X1)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X1] : (sP14(X1) <=> (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f27, plain, ! [X1] : (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13)))) | ~ sP15(X1)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X1] : (sP15(X1) <=> (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f28, plain, ! [X1] : (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11)))) | ~ sP16(X1)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X1] : (sP16(X1) <=> (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f29, plain, ! [X1] : (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9)))) | ~ sP17(X1)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X1] : (sP17(X1) <=> (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f30, plain, ! [X1] : (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7)))) | ~ sP18(X1)), inference(usedef, [], [e30])).
fof(e30, plain, ! [X1] : (sP18(X1) <=> (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f31, plain, ! [X1] : (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5)))) | ~ sP19(X1)), inference(usedef, [], [e31])).
fof(e31, plain, ! [X1] : (sP19(X1) <=> (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f32, plain, ! [X1] : (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3)))) | ~ sP20(X1)), inference(usedef, [], [e32])).
fof(e32, plain, ! [X1] : (sP20(X1) <=> (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f33, plain, ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & sP20(X1) & sP19(X1) & sP18(X1) & sP17(X1) & sP16(X1) & sP15(X1) & sP14(X1) & sP13(X1) & sP12(X1) & sP11(X1) & sP10(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP4(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP9(X1)) | ~ sP21(X1)), inference(usedef, [], [e33])).
fof(e33, plain, ! [X1] : (sP21(X1) <=> ((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & sP20(X1) & sP19(X1) & sP18(X1) & sP17(X1) & sP16(X1) & sP15(X1) & sP14(X1) & sP13(X1) & sP12(X1) & sP11(X1) & sP10(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP4(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP9(X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f11, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p100(X1) | p101(X1) | (? [X24] : (p101(X24) & ~ p102(X24) & p2(X24) & r1(X1, X24)) & ? [X25] : (p101(X25) & ~ p102(X25) & ~ p2(X25) & r1(X1, X25)))) & (~ p101(X1) | p102(X1) | (? [X26] : (p102(X26) & ~ p103(X26) & p3(X26) & r1(X1, X26)) & ? [X27] : (p102(X27) & ~ p103(X27) & ~ p3(X27) & r1(X1, X27)))) & (~ p102(X1) | p103(X1) | (? [X28] : (p103(X28) & ~ p104(X28) & p4(X28) & r1(X1, X28)) & ? [X29] : (p103(X29) & ~ p104(X29) & ~ p4(X29) & r1(X1, X29)))) & (~ p103(X1) | p104(X1) | (? [X30] : (p104(X30) & ~ p105(X30) & p5(X30) & r1(X1, X30)) & ? [X31] : (p104(X31) & ~ p105(X31) & ~ p5(X31) & r1(X1, X31)))) & (~ p104(X1) | p105(X1) | (? [X32] : (p105(X32) & ~ p106(X32) & p6(X32) & r1(X1, X32)) & ? [X33] : (p105(X33) & ~ p106(X33) & ~ p6(X33) & r1(X1, X33)))) & (~ p105(X1) | p106(X1) | (? [X34] : (p106(X34) & ~ p107(X34) & p7(X34) & r1(X1, X34)) & ? [X35] : (p106(X35) & ~ p107(X35) & ~ p7(X35) & r1(X1, X35)))) & (~ p106(X1) | p107(X1) | (? [X36] : (p107(X36) & ~ p108(X36) & p8(X36) & r1(X1, X36)) & ? [X37] : (p107(X37) & ~ p108(X37) & ~ p8(X37) & r1(X1, X37)))) & (~ p107(X1) | p108(X1) | (? [X38] : (p108(X38) & ~ p109(X38) & p9(X38) & r1(X1, X38)) & ? [X39] : (p108(X39) & ~ p109(X39) & ~ p9(X39) & r1(X1, X39)))) & (~ p108(X1) | p109(X1) | (? [X40] : (p109(X40) & ~ p110(X40) & p10(X40) & r1(X1, X40)) & ? [X41] : (p109(X41) & ~ p110(X41) & ~ p10(X41) & r1(X1, X41)))) & (~ p109(X1) | p110(X1) | (? [X42] : (p110(X42) & p11(X42) & r1(X1, X42)) & ? [X43] : (p110(X43) & ~ p11(X43) & r1(X1, X43))))) | ~ r1(X0, X1)) & ! [X44] : (p5(X44) | ~ r1(X0, X44))), inference(flattening, [], [f10])).
fof(f10, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & ((~ p100(X1) | p101(X1)) | (? [X24] : ((p101(X24) & ~ p102(X24) & p2(X24)) & r1(X1, X24)) & ? [X25] : ((p101(X25) & ~ p102(X25) & ~ p2(X25)) & r1(X1, X25)))) & ((~ p101(X1) | p102(X1)) | (? [X26] : ((p102(X26) & ~ p103(X26) & p3(X26)) & r1(X1, X26)) & ? [X27] : ((p102(X27) & ~ p103(X27) & ~ p3(X27)) & r1(X1, X27)))) & ((~ p102(X1) | p103(X1)) | (? [X28] : ((p103(X28) & ~ p104(X28) & p4(X28)) & r1(X1, X28)) & ? [X29] : ((p103(X29) & ~ p104(X29) & ~ p4(X29)) & r1(X1, X29)))) & ((~ p103(X1) | p104(X1)) | (? [X30] : ((p104(X30) & ~ p105(X30) & p5(X30)) & r1(X1, X30)) & ? [X31] : ((p104(X31) & ~ p105(X31) & ~ p5(X31)) & r1(X1, X31)))) & ((~ p104(X1) | p105(X1)) | (? [X32] : ((p105(X32) & ~ p106(X32) & p6(X32)) & r1(X1, X32)) & ? [X33] : ((p105(X33) & ~ p106(X33) & ~ p6(X33)) & r1(X1, X33)))) & ((~ p105(X1) | p106(X1)) | (? [X34] : ((p106(X34) & ~ p107(X34) & p7(X34)) & r1(X1, X34)) & ? [X35] : ((p106(X35) & ~ p107(X35) & ~ p7(X35)) & r1(X1, X35)))) & ((~ p106(X1) | p107(X1)) | (? [X36] : ((p107(X36) & ~ p108(X36) & p8(X36)) & r1(X1, X36)) & ? [X37] : ((p107(X37) & ~ p108(X37) & ~ p8(X37)) & r1(X1, X37)))) & ((~ p107(X1) | p108(X1)) | (? [X38] : ((p108(X38) & ~ p109(X38) & p9(X38)) & r1(X1, X38)) & ? [X39] : ((p108(X39) & ~ p109(X39) & ~ p9(X39)) & r1(X1, X39)))) & ((~ p108(X1) | p109(X1)) | (? [X40] : ((p109(X40) & ~ p110(X40) & p10(X40)) & r1(X1, X40)) & ? [X41] : ((p109(X41) & ~ p110(X41) & ~ p10(X41)) & r1(X1, X41)))) & ((~ p109(X1) | p110(X1)) | (? [X42] : ((p110(X42) & p11(X42)) & r1(X1, X42)) & ? [X43] : ((p110(X43) & ~ p11(X43)) & r1(X1, X43))))) | ~ r1(X0, X1))) & ! [X44] : (p5(X44) | ~ r1(X0, X44))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X24] : (~ (p101(X24) & ~ p102(X24) & p2(X24)) | ~ r1(X1, X24)) & ~ ! [X25] : (~ (p101(X25) & ~ p102(X25) & ~ p2(X25)) | ~ r1(X1, X25)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X26] : (~ (p102(X26) & ~ p103(X26) & p3(X26)) | ~ r1(X1, X26)) & ~ ! [X27] : (~ (p102(X27) & ~ p103(X27) & ~ p3(X27)) | ~ r1(X1, X27)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X28] : (~ (p103(X28) & ~ p104(X28) & p4(X28)) | ~ r1(X1, X28)) & ~ ! [X29] : (~ (p103(X29) & ~ p104(X29) & ~ p4(X29)) | ~ r1(X1, X29)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X30] : (~ (p104(X30) & ~ p105(X30) & p5(X30)) | ~ r1(X1, X30)) & ~ ! [X31] : (~ (p104(X31) & ~ p105(X31) & ~ p5(X31)) | ~ r1(X1, X31)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X32] : (~ (p105(X32) & ~ p106(X32) & p6(X32)) | ~ r1(X1, X32)) & ~ ! [X33] : (~ (p105(X33) & ~ p106(X33) & ~ p6(X33)) | ~ r1(X1, X33)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X34] : (~ (p106(X34) & ~ p107(X34) & p7(X34)) | ~ r1(X1, X34)) & ~ ! [X35] : (~ (p106(X35) & ~ p107(X35) & ~ p7(X35)) | ~ r1(X1, X35)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X36] : (~ (p107(X36) & ~ p108(X36) & p8(X36)) | ~ r1(X1, X36)) & ~ ! [X37] : (~ (p107(X37) & ~ p108(X37) & ~ p8(X37)) | ~ r1(X1, X37)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X38] : (~ (p108(X38) & ~ p109(X38) & p9(X38)) | ~ r1(X1, X38)) & ~ ! [X39] : (~ (p108(X39) & ~ p109(X39) & ~ p9(X39)) | ~ r1(X1, X39)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X40] : (~ (p109(X40) & ~ p110(X40) & p10(X40)) | ~ r1(X1, X40)) & ~ ! [X41] : (~ (p109(X41) & ~ p110(X41) & ~ p10(X41)) | ~ r1(X1, X41)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X42] : (~ (p110(X42) & p11(X42)) | ~ r1(X1, X42)) & ~ ! [X43] : (~ (p110(X43) & ~ p11(X43)) | ~ r1(X1, X43))))) | ~ r1(X0, X1))) | ~ ! [X44] : (p5(X44) | ~ r1(X0, X44))), inference(pure_predicate_removal, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X24] : (~ (p101(X24) & ~ p102(X24) & p2(X24)) | ~ r1(X1, X24)) & ~ ! [X25] : (~ (p101(X25) & ~ p102(X25) & ~ p2(X25)) | ~ r1(X1, X25)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X26] : (~ (p102(X26) & ~ p103(X26) & p3(X26)) | ~ r1(X1, X26)) & ~ ! [X27] : (~ (p102(X27) & ~ p103(X27) & ~ p3(X27)) | ~ r1(X1, X27)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X28] : (~ (p103(X28) & ~ p104(X28) & p4(X28)) | ~ r1(X1, X28)) & ~ ! [X29] : (~ (p103(X29) & ~ p104(X29) & ~ p4(X29)) | ~ r1(X1, X29)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X30] : (~ (p104(X30) & ~ p105(X30) & p5(X30)) | ~ r1(X1, X30)) & ~ ! [X31] : (~ (p104(X31) & ~ p105(X31) & ~ p5(X31)) | ~ r1(X1, X31)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X32] : (~ (p105(X32) & ~ p106(X32) & p6(X32)) | ~ r1(X1, X32)) & ~ ! [X33] : (~ (p105(X33) & ~ p106(X33) & ~ p6(X33)) | ~ r1(X1, X33)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X34] : (~ (p106(X34) & ~ p107(X34) & p7(X34)) | ~ r1(X1, X34)) & ~ ! [X35] : (~ (p106(X35) & ~ p107(X35) & ~ p7(X35)) | ~ r1(X1, X35)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X36] : (~ (p107(X36) & ~ p108(X36) & p8(X36)) | ~ r1(X1, X36)) & ~ ! [X37] : (~ (p107(X37) & ~ p108(X37) & ~ p8(X37)) | ~ r1(X1, X37)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X38] : (~ (p108(X38) & ~ p109(X38) & p9(X38)) | ~ r1(X1, X38)) & ~ ! [X39] : (~ (p108(X39) & ~ p109(X39) & ~ p9(X39)) | ~ r1(X1, X39)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X40] : (~ (p109(X40) & ~ p110(X40) & p10(X40)) | ~ r1(X1, X40)) & ~ ! [X41] : (~ (p109(X41) & ~ p110(X41) & ~ p10(X41)) | ~ r1(X1, X41)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X42] : (~ (p110(X42) & ~ p111(X42) & p11(X42)) | ~ r1(X1, X42)) & ~ ! [X43] : (~ (p110(X43) & ~ p111(X43) & ~ p11(X43)) | ~ r1(X1, X43))))) | ~ r1(X0, X1))) | ~ ! [X44] : (p5(X44) | ~ r1(X0, X44))), inference(flattening, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X24] : (~ (p101(X24) & ~ p102(X24) & p2(X24)) | ~ r1(X1, X24)) & ~ ! [X25] : (~ (p101(X25) & ~ p102(X25) & ~ p2(X25)) | ~ r1(X1, X25)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X26] : (~ (p102(X26) & ~ p103(X26) & p3(X26)) | ~ r1(X1, X26)) & ~ ! [X27] : (~ (p102(X27) & ~ p103(X27) & ~ p3(X27)) | ~ r1(X1, X27)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X28] : (~ (p103(X28) & ~ p104(X28) & p4(X28)) | ~ r1(X1, X28)) & ~ ! [X29] : (~ (p103(X29) & ~ p104(X29) & ~ p4(X29)) | ~ r1(X1, X29)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X30] : (~ (p104(X30) & ~ p105(X30) & p5(X30)) | ~ r1(X1, X30)) & ~ ! [X31] : (~ (p104(X31) & ~ p105(X31) & ~ p5(X31)) | ~ r1(X1, X31)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X32] : (~ (p105(X32) & ~ p106(X32) & p6(X32)) | ~ r1(X1, X32)) & ~ ! [X33] : (~ (p105(X33) & ~ p106(X33) & ~ p6(X33)) | ~ r1(X1, X33)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X34] : (~ (p106(X34) & ~ p107(X34) & p7(X34)) | ~ r1(X1, X34)) & ~ ! [X35] : (~ (p106(X35) & ~ p107(X35) & ~ p7(X35)) | ~ r1(X1, X35)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X36] : (~ (p107(X36) & ~ p108(X36) & p8(X36)) | ~ r1(X1, X36)) & ~ ! [X37] : (~ (p107(X37) & ~ p108(X37) & ~ p8(X37)) | ~ r1(X1, X37)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X38] : (~ (p108(X38) & ~ p109(X38) & p9(X38)) | ~ r1(X1, X38)) & ~ ! [X39] : (~ (p108(X39) & ~ p109(X39) & ~ p9(X39)) | ~ r1(X1, X39)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X40] : (~ (p109(X40) & ~ p110(X40) & p10(X40)) | ~ r1(X1, X40)) & ~ ! [X41] : (~ (p109(X41) & ~ p110(X41) & ~ p10(X41)) | ~ r1(X1, X41)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X42] : (~ (p110(X42) & ~ p111(X42) & p11(X42)) | ~ r1(X1, X42)) & ~ ! [X43] : (~ (p110(X43) & ~ p111(X43) & ~ p11(X43)) | ~ r1(X1, X43))))) | ~ r1(X0, X1))) | ~ ! [X44] : (p5(X44) | ~ r1(X0, X44))), inference(rectify, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ p106(X1) | ((~ p7(X1) | ! [X0] : (~ p106(X0) | p7(X0) | ~ r1(X1, X0))) & (p7(X1) | ! [X0] : (~ p106(X0) | ~ p7(X0) | ~ r1(X1, X0))))) & (~ p107(X1) | ((~ p8(X1) | ! [X0] : (~ p107(X0) | p8(X0) | ~ r1(X1, X0))) & (p8(X1) | ! [X0] : (~ p107(X0) | ~ p8(X0) | ~ r1(X1, X0))))) & (~ p108(X1) | ((~ p9(X1) | ! [X0] : (~ p108(X0) | p9(X0) | ~ r1(X1, X0))) & (p9(X1) | ! [X0] : (~ p108(X0) | ~ p9(X0) | ~ r1(X1, X0))))) & (~ p109(X1) | ((~ p10(X1) | ! [X0] : (~ p109(X0) | p10(X0) | ~ r1(X1, X0))) & (p10(X1) | ! [X0] : (~ p109(X0) | ~ p10(X0) | ~ r1(X1, X0))))) & (~ p110(X1) | ((~ p11(X1) | ! [X0] : (~ p110(X0) | p11(X0) | ~ r1(X1, X0))) & (p11(X1) | ! [X0] : (~ p110(X0) | ~ p11(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & p7(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & ~ p7(X0)) | ~ r1(X1, X0)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & p8(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & ~ p8(X0)) | ~ r1(X1, X0)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & p9(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & ~ p9(X0)) | ~ r1(X1, X0)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & p10(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & ~ p10(X0)) | ~ r1(X1, X0)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & p11(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & ~ p11(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p5(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ p106(X1) | ((~ p7(X1) | ! [X0] : (~ p106(X0) | p7(X0) | ~ r1(X1, X0))) & (p7(X1) | ! [X0] : (~ p106(X0) | ~ p7(X0) | ~ r1(X1, X0))))) & (~ p107(X1) | ((~ p8(X1) | ! [X0] : (~ p107(X0) | p8(X0) | ~ r1(X1, X0))) & (p8(X1) | ! [X0] : (~ p107(X0) | ~ p8(X0) | ~ r1(X1, X0))))) & (~ p108(X1) | ((~ p9(X1) | ! [X0] : (~ p108(X0) | p9(X0) | ~ r1(X1, X0))) & (p9(X1) | ! [X0] : (~ p108(X0) | ~ p9(X0) | ~ r1(X1, X0))))) & (~ p109(X1) | ((~ p10(X1) | ! [X0] : (~ p109(X0) | p10(X0) | ~ r1(X1, X0))) & (p10(X1) | ! [X0] : (~ p109(X0) | ~ p10(X0) | ~ r1(X1, X0))))) & (~ p110(X1) | ((~ p11(X1) | ! [X0] : (~ p110(X0) | p11(X0) | ~ r1(X1, X0))) & (p11(X1) | ! [X0] : (~ p110(X0) | ~ p11(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & p7(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & ~ p7(X0)) | ~ r1(X1, X0)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & p8(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & ~ p8(X0)) | ~ r1(X1, X0)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & p9(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & ~ p9(X0)) | ~ r1(X1, X0)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & p10(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & ~ p10(X0)) | ~ r1(X1, X0)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & p11(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & ~ p11(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p5(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.010.p', main)).
fof(f6756, plain, (p101(sK42) | ~ r1(sK42, sK42) | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f6755, f260])).
fof(f260, plain, sP8(sK42), inference(resolution, [], [f250, f123])).
fof(f123, plain, ! [X0] : (~ sP21(X0) | sP8(X0)), inference(cnf_transformation, [], [f36])).
fof(f36, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p107(X0) | p106(X0)) & (~ p108(X0) | p107(X0)) & (~ p109(X0) | p108(X0)) & (~ p110(X0) | p109(X0)) & sP20(X0) & sP19(X0) & sP18(X0) & sP17(X0) & sP16(X0) & sP15(X0) & sP14(X0) & sP13(X0) & sP12(X0) & sP11(X0) & sP10(X0) & sP8(X0) & sP7(X0) & sP6(X0) & sP5(X0) & sP4(X0) & sP3(X0) & sP2(X0) & sP1(X0) & sP0(X0) & sP9(X0)) | ~ sP21(X0)), inference(rectify, [], [f35])).
fof(f35, plain, ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & sP20(X1) & sP19(X1) & sP18(X1) & sP17(X1) & sP16(X1) & sP15(X1) & sP14(X1) & sP13(X1) & sP12(X1) & sP11(X1) & sP10(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP4(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP9(X1)) | ~ sP21(X1)), inference(nnf_transformation, [], [f33])).
fof(f250, plain, sP21(sK42), inference(resolution, [], [f246, f112])).
fof(f246, plain, ! [X1] : (~ r1(sK42, X1) | sP21(X1)), inference(cnf_transformation, [], [f111])).
fof(f6755, plain, (~ sP8(sK42) | p101(sK42) | ~ r1(sK42, sK42) | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f6750, f248])).
fof(f248, plain, p100(sK42), inference(cnf_transformation, [], [f111])).
fof(f6750, plain, (~ p100(sK42) | ~ sP8(sK42) | p101(sK42) | ~ r1(sK42, sK42) | (spl43_29 | ~ spl43_30)), inference(resolution, [], [f6739, f431])).
fof(f431, plain, ! [X0, X1] : (r1(X1, sK24(X0)) | ~ p100(X0) | ~ sP8(X0) | p101(X0) | ~ r1(X1, X0)), inference(resolution, [], [f177, f113])).
fof(f113, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | r1(X0, X2) | ~ r1(X0, X1)), inference(cnf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1, X2] : (r1(X0, X2) | ~ r1(X1, X2) | ~ r1(X0, X1)), inference(flattening, [], [f8])).
fof(f8, plain, ! [X0, X1, X2] : (r1(X0, X2) | (~ r1(X1, X2) | ~ r1(X0, X1))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1, X2] : ((r1(X1, X2) & r1(X0, X1)) => r1(X0, X2)), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.010.p', transitivity)).
fof(f177, plain, ! [X0] : (r1(X0, sK24(X0)) | p101(X0) | ~ p100(X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f68])).
fof(f68, plain, ! [X0] : (~ p100(X0) | p101(X0) | ((p101(sK24(X0)) & ~ p102(sK24(X0)) & p2(sK24(X0)) & r1(X0, sK24(X0))) & (p101(sK25(X0)) & ~ p102(sK25(X0)) & ~ p2(sK25(X0)) & r1(X0, sK25(X0)))) | ~ sP8(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK24, sK25])], [f65, f67, f66])).
fof(f66, plain, ! [X0] : (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) => (p101(sK24(X0)) & ~ p102(sK24(X0)) & p2(sK24(X0)) & r1(X0, sK24(X0)))), introduced(choice_axiom, [])).
fof(f67, plain, ! [X0] : (? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2)) => (p101(sK25(X0)) & ~ p102(sK25(X0)) & ~ p2(sK25(X0)) & r1(X0, sK25(X0)))), introduced(choice_axiom, [])).
fof(f65, plain, ! [X0] : (~ p100(X0) | p101(X0) | (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) & ? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2))) | ~ sP8(X0)), inference(rectify, [], [f64])).
fof(f64, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X24] : (p101(X24) & ~ p102(X24) & p2(X24) & r1(X1, X24)) & ? [X25] : (p101(X25) & ~ p102(X25) & ~ p2(X25) & r1(X1, X25))) | ~ sP8(X1)), inference(nnf_transformation, [], [f20])).
fof(f6739, plain, (~ r1(sK42, sK24(sK42)) | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f6738, f519])).
fof(f519, plain, (p101(sK24(sK42)) | ~ spl43_30), inference(avatar_component_clause, [], [f517])).
fof(f517, plain, (spl43_30 <=> p101(sK24(sK42))), introduced(avatar_definition, [new_symbols(naming, [spl43_30])])).
fof(f6738, plain, (~ r1(sK42, sK24(sK42)) | ~ p101(sK24(sK42)) | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f6737, f513])).
fof(f513, plain, (~ p102(sK24(sK42)) | spl43_29), inference(avatar_component_clause, [], [f512])).
fof(f512, plain, (spl43_29 <=> p102(sK24(sK42))), introduced(avatar_definition, [new_symbols(naming, [spl43_29])])).
fof(f6737, plain, (p102(sK24(sK42)) | ~ r1(sK42, sK24(sK42)) | ~ p101(sK24(sK42)) | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f6731, f449])).
fof(f449, plain, sP7(sK24(sK42)), inference(resolution, [], [f434, f122])).
fof(f122, plain, ! [X0] : (~ sP21(X0) | sP7(X0)), inference(cnf_transformation, [], [f36])).
fof(f434, plain, sP21(sK24(sK42)), inference(subsumption_resolution, [], [f433, f260])).
fof(f433, plain, (~ sP8(sK42) | sP21(sK24(sK42))), inference(subsumption_resolution, [], [f432, f248])).
fof(f432, plain, (~ p100(sK42) | ~ sP8(sK42) | sP21(sK24(sK42))), inference(subsumption_resolution, [], [f429, f247])).
fof(f429, plain, (p101(sK42) | ~ p100(sK42) | ~ sP8(sK42) | sP21(sK24(sK42))), inference(resolution, [], [f177, f246])).
fof(f6731, plain, (~ sP7(sK24(sK42)) | p102(sK24(sK42)) | ~ r1(sK42, sK24(sK42)) | ~ p101(sK24(sK42)) | (spl43_29 | ~ spl43_30)), inference(resolution, [], [f6730, f2574])).
fof(f2574, plain, (sP6(sK26(sK24(sK42))) | (spl43_29 | ~ spl43_30)), inference(resolution, [], [f2556, f121])).
fof(f121, plain, ! [X0] : (~ sP21(X0) | sP6(X0)), inference(cnf_transformation, [], [f36])).
fof(f2556, plain, (sP21(sK26(sK24(sK42))) | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f2555, f260])).
fof(f2555, plain, (sP21(sK26(sK24(sK42))) | ~ sP8(sK42) | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f2554, f248])).
fof(f2554, plain, (sP21(sK26(sK24(sK42))) | ~ p100(sK42) | ~ sP8(sK42) | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f2553, f247])).
fof(f2553, plain, (sP21(sK26(sK24(sK42))) | p101(sK42) | ~ p100(sK42) | ~ sP8(sK42) | (spl43_29 | ~ spl43_30)), inference(subsumption_resolution, [], [f2552, f519])).
fof(f2552, plain, (~ p101(sK24(sK42)) | sP21(sK26(sK24(sK42))) | p101(sK42) | ~ p100(sK42) | ~ sP8(sK42) | spl43_29), inference(subsumption_resolution, [], [f2551, f513])).
fof(f2551, plain, (p102(sK24(sK42)) | ~ p101(sK24(sK42)) | sP21(sK26(sK24(sK42))) | p101(sK42) | ~ p100(sK42) | ~ sP8(sK42)), inference(subsumption_resolution, [], [f2515, f449])).
fof(f2515, plain, (~ sP7(sK24(sK42)) | p102(sK24(sK42)) | ~ p101(sK24(sK42)) | sP21(sK26(sK24(sK42))) | p101(sK42) | ~ p100(sK42) | ~ sP8(sK42)), inference(resolution, [], [f1237, f177])).
fof(f1237, plain, ! [X0] : (~ r1(sK42, X0) | ~ sP7(X0) | p102(X0) | ~ p101(X0) | sP21(sK26(X0))), inference(resolution, [], [f528, f246])).
fof(f528, plain, ! [X0, X1] : (r1(X1, sK26(X0)) | ~ p101(X0) | ~ sP7(X0) | p102(X0) | ~ r1(X1, X0)), inference(resolution, [], [f185, f113])).
fof(f185, plain, ! [X0] : (r1(X0, sK26(X0)) | p102(X0) | ~ p101(X0) | ~ sP7(X0)), inference(cnf_transformation, [], [f73])).
fof(f73, plain, ! [X0] : (~ p101(X0) | p102(X0) | ((p102(sK26(X0)) & ~ p103(sK26(X0)) & p3(sK26(X0)) & r1(X0, sK26(X0))) & (p102(sK27(X0)) & ~ p103(sK27(X0)) & ~ p3(sK27(X0)) & r1(X0, sK27(X0)))) | ~ sP7(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK26, sK27])], [f70, f72, f71])).
fof(f71, plain, ! [X0] : (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) => (p102(sK26(X0)) & ~ p103(sK26(X0)) & p3(sK26(X0)) & r1(X0, sK26(X0)))), introduced(choice_axiom, [])).
fof(f72, plain, ! [X0] : (? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2)) => (p102(sK27(X0)) & ~ p103(sK27(X0)) & ~ p3(sK27(X0)) & r1(X0, sK27(X0)))), introduced(choice_axiom, [])).
fof(f70, plain, ! [X0] : (~ p101(X0) | p102(X0) | (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) & ? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2))) | ~ sP7(X0)), inference(rectify, [], [f69])).
fof(f69, plain, ! [X1] : (~ p101(X1) | p102(X1) | (? [X26] : (p102(X26) & ~ p103(X26) & p3(X26) & r1(X1, X26)) & ? [X27] : (p102(X27) & ~ p103(X27) & ~ p3(X27) & r1(X1, X27))) | ~ sP7(X1)), inference(nnf_transformation, [], [f19])).
fof(f6730, plain, ! [X3] : (~ sP6(sK26(X3)) | ~ sP7(X3) | p102(X3) | ~ r1(sK42, X3) | ~ p101(X3)), inference(subsumption_resolution, [], [f6729, f528])).
fof(f6729, plain, ! [X3] : (~ p101(X3) | ~ sP7(X3) | p102(X3) | ~ r1(sK42, X3) | ~ sP6(sK26(X3)) | ~ r1(sK42, sK26(X3))), inference(subsumption_resolution, [], [f6728, f187])).
fof(f187, plain, ! [X0] : (~ p103(sK26(X0)) | p102(X0) | ~ p101(X0) | ~ sP7(X0)), inference(cnf_transformation, [], [f73])).
fof(f6728, plain, ! [X3] : (~ p101(X3) | ~ sP7(X3) | p102(X3) | ~ r1(sK42, X3) | ~ sP6(sK26(X3)) | p103(sK26(X3)) | ~ r1(sK42, sK26(X3))), inference(subsumption_resolution, [], [f6721, f188])).
fof(f188, plain, ! [X0] : (p102(sK26(X0)) | p102(X0) | ~ p101(X0) | ~ sP7(X0)), inference(cnf_transformation, [], [f73])).
fof(f6721, plain, ! [X3] : (~ p101(X3) | ~ sP7(X3) | p102(X3) | ~ r1(sK42, X3) | ~ sP6(sK26(X3)) | ~ p102(sK26(X3)) | p103(sK26(X3)) | ~ r1(sK42, sK26(X3))), inference(duplicate_literal_removal, [], [f6720])).
fof(f6720, plain, ! [X3] : (~ p101(X3) | ~ sP7(X3) | p102(X3) | ~ r1(sK42, X3) | ~ sP6(sK26(X3)) | ~ p102(sK26(X3)) | ~ sP6(sK26(X3)) | p103(sK26(X3)) | ~ r1(sK42, sK26(X3))), inference(resolution, [], [f5927, f1715])).
fof(f1715, plain, ! [X7] : (~ sP5(sK29(X7)) | ~ p102(X7) | ~ sP6(X7) | p103(X7) | ~ r1(sK42, X7)), inference(subsumption_resolution, [], [f1714, f192])).
fof(f192, plain, ! [X0] : (p103(sK29(X0)) | p103(X0) | ~ p102(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f78])).
fof(f78, plain, ! [X0] : (~ p102(X0) | p103(X0) | ((p103(sK28(X0)) & ~ p104(sK28(X0)) & p4(sK28(X0)) & r1(X0, sK28(X0))) & (p103(sK29(X0)) & ~ p104(sK29(X0)) & ~ p4(sK29(X0)) & r1(X0, sK29(X0)))) | ~ sP6(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK28, sK29])], [f75, f77, f76])).
fof(f76, plain, ! [X0] : (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) => (p103(sK28(X0)) & ~ p104(sK28(X0)) & p4(sK28(X0)) & r1(X0, sK28(X0)))), introduced(choice_axiom, [])).
fof(f77, plain, ! [X0] : (? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2)) => (p103(sK29(X0)) & ~ p104(sK29(X0)) & ~ p4(sK29(X0)) & r1(X0, sK29(X0)))), introduced(choice_axiom, [])).
fof(f75, plain, ! [X0] : (~ p102(X0) | p103(X0) | (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) & ? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2))) | ~ sP6(X0)), inference(rectify, [], [f74])).
fof(f74, plain, ! [X1] : (~ p102(X1) | p103(X1) | (? [X28] : (p103(X28) & ~ p104(X28) & p4(X28) & r1(X1, X28)) & ? [X29] : (p103(X29) & ~ p104(X29) & ~ p4(X29) & r1(X1, X29))) | ~ sP6(X1)), inference(nnf_transformation, [], [f18])).
fof(f1714, plain, ! [X7] : (~ sP5(sK29(X7)) | ~ p103(sK29(X7)) | ~ p102(X7) | ~ sP6(X7) | p103(X7) | ~ r1(sK42, X7)), inference(subsumption_resolution, [], [f1687, f191])).
fof(f191, plain, ! [X0] : (~ p104(sK29(X0)) | p103(X0) | ~ p102(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f78])).
fof(f1687, plain, ! [X7] : (~ sP5(sK29(X7)) | p104(sK29(X7)) | ~ p103(sK29(X7)) | ~ p102(X7) | ~ sP6(X7) | p103(X7) | ~ r1(sK42, X7)), inference(resolution, [], [f1345, f531])).
fof(f531, plain, ! [X0, X1] : (r1(X1, sK29(X0)) | ~ p102(X0) | ~ sP6(X0) | p103(X0) | ~ r1(X1, X0)), inference(resolution, [], [f189, f113])).
fof(f189, plain, ! [X0] : (r1(X0, sK29(X0)) | p103(X0) | ~ p102(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f78])).
fof(f1345, plain, ! [X1] : (~ r1(sK42, X1) | ~ sP5(X1) | p104(X1) | ~ p103(X1)), inference(subsumption_resolution, [], [f1321, f198])).
fof(f198, plain, ! [X0] : (~ p5(sK31(X0)) | p104(X0) | ~ p103(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f83])).
fof(f83, plain, ! [X0] : (~ p103(X0) | p104(X0) | ((p104(sK30(X0)) & ~ p105(sK30(X0)) & p5(sK30(X0)) & r1(X0, sK30(X0))) & (p104(sK31(X0)) & ~ p105(sK31(X0)) & ~ p5(sK31(X0)) & r1(X0, sK31(X0)))) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK30, sK31])], [f80, f82, f81])).
fof(f81, plain, ! [X0] : (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) => (p104(sK30(X0)) & ~ p105(sK30(X0)) & p5(sK30(X0)) & r1(X0, sK30(X0)))), introduced(choice_axiom, [])).
fof(f82, plain, ! [X0] : (? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2)) => (p104(sK31(X0)) & ~ p105(sK31(X0)) & ~ p5(sK31(X0)) & r1(X0, sK31(X0)))), introduced(choice_axiom, [])).
fof(f80, plain, ! [X0] : (~ p103(X0) | p104(X0) | (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) & ? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2))) | ~ sP5(X0)), inference(rectify, [], [f79])).
fof(f79, plain, ! [X1] : (~ p103(X1) | p104(X1) | (? [X30] : (p104(X30) & ~ p105(X30) & p5(X30) & r1(X1, X30)) & ? [X31] : (p104(X31) & ~ p105(X31) & ~ p5(X31) & r1(X1, X31))) | ~ sP5(X1)), inference(nnf_transformation, [], [f17])).
fof(f1321, plain, ! [X1] : (~ p103(X1) | ~ sP5(X1) | p104(X1) | ~ r1(sK42, X1) | p5(sK31(X1))), inference(resolution, [], [f537, f245])).
fof(f245, plain, ! [X2] : (~ r1(sK42, X2) | p5(X2)), inference(cnf_transformation, [], [f111])).
fof(f537, plain, ! [X0, X1] : (r1(X1, sK31(X0)) | ~ p103(X0) | ~ sP5(X0) | p104(X0) | ~ r1(X1, X0)), inference(resolution, [], [f197, f113])).
fof(f197, plain, ! [X0] : (r1(X0, sK31(X0)) | p104(X0) | ~ p103(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f83])).
fof(f5927, plain, ! [X6] : (sP5(sK29(sK26(X6))) | ~ p101(X6) | ~ sP7(X6) | p102(X6) | ~ r1(sK42, X6) | ~ sP6(sK26(X6))), inference(resolution, [], [f2856, f120])).
fof(f120, plain, ! [X0] : (~ sP21(X0) | sP5(X0)), inference(cnf_transformation, [], [f36])).
fof(f2856, plain, ! [X4] : (sP21(sK29(sK26(X4))) | ~ sP6(sK26(X4)) | ~ p101(X4) | ~ sP7(X4) | p102(X4) | ~ r1(sK42, X4)), inference(subsumption_resolution, [], [f2855, f188])).
fof(f2855, plain, ! [X4] : (~ sP6(sK26(X4)) | ~ p102(sK26(X4)) | sP21(sK29(sK26(X4))) | ~ p101(X4) | ~ sP7(X4) | p102(X4) | ~ r1(sK42, X4)), inference(subsumption_resolution, [], [f2824, f187])).
fof(f2824, plain, ! [X4] : (~ sP6(sK26(X4)) | p103(sK26(X4)) | ~ p102(sK26(X4)) | sP21(sK29(sK26(X4))) | ~ p101(X4) | ~ sP7(X4) | p102(X4) | ~ r1(sK42, X4)), inference(resolution, [], [f1265, f528])).
fof(f1265, plain, ! [X0] : (~ r1(sK42, X0) | ~ sP6(X0) | p103(X0) | ~ p102(X0) | sP21(sK29(X0))), inference(resolution, [], [f531, f246])).
fof(f2308, plain, ~ spl43_29, inference(avatar_contradiction_clause, [], [f2307])).
fof(f2307, plain, ($false | ~ spl43_29), inference(subsumption_resolution, [], [f2306, f260])).
fof(f2306, plain, (~ sP8(sK42) | ~ spl43_29), inference(subsumption_resolution, [], [f2305, f248])).
fof(f2305, plain, (~ p100(sK42) | ~ sP8(sK42) | ~ spl43_29), inference(subsumption_resolution, [], [f2304, f247])).
fof(f2304, plain, (p101(sK42) | ~ p100(sK42) | ~ sP8(sK42) | ~ spl43_29), inference(resolution, [], [f514, f179])).
fof(f179, plain, ! [X0] : (~ p102(sK24(X0)) | p101(X0) | ~ p100(X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f68])).
fof(f514, plain, (p102(sK24(sK42)) | ~ spl43_29), inference(avatar_component_clause, [], [f512])).
fof(f1669, plain, spl43_30, inference(avatar_contradiction_clause, [], [f1668])).
fof(f1668, plain, ($false | spl43_30), inference(subsumption_resolution, [], [f1667, f260])).
fof(f1667, plain, (~ sP8(sK42) | spl43_30), inference(subsumption_resolution, [], [f1666, f248])).
fof(f1666, plain, (~ p100(sK42) | ~ sP8(sK42) | spl43_30), inference(subsumption_resolution, [], [f1665, f247])).
fof(f1665, plain, (p101(sK42) | ~ p100(sK42) | ~ sP8(sK42) | spl43_30), inference(resolution, [], [f518, f180])).
fof(f180, plain, ! [X0] : (p101(sK24(X0)) | p101(X0) | ~ p100(X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f68])).
fof(f518, plain, (~ p101(sK24(sK42)) | spl43_30), inference(avatar_component_clause, [], [f517])).