fof(f21282, plain, $false, inference(avatar_sat_refutation, [], [f1247, f6583, f6733, f6816, f6817, f7386, f9799, f9879, f9979, f9980, f9981, f19666, f19671, f19676, f19739, f19767, f19772, f21157, f21216, f21281])).
fof(f21281, plain, ~ spl121_188, inference(avatar_contradiction_clause, [], [f21280])).
fof(f21280, plain, ($false | ~ spl121_188), inference(subsumption_resolution, [], [f21279, f709])).
fof(f709, plain, sP50(sK111), inference(resolution, [], [f704, f280])).
fof(f280, plain, ! [X0] : (~ sP59(X0) | sP50(X0)), inference(cnf_transformation, [], [f70])).
fof(f70, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & sP58(X0) & sP57(X0) & sP56(X0) & sP55(X0) & sP54(X0) & sP53(X0) & sP51(X0) & sP50(X0) & sP49(X0) & sP48(X0) & sP52(X0)) | ~ sP59(X0)), inference(rectify, [], [f69])).
fof(f69, plain, ! [X23] : (((~ p101(X23) | p100(X23)) & (~ p102(X23) | p101(X23)) & (~ p103(X23) | p102(X23)) & (~ p104(X23) | p103(X23)) & (~ p105(X23) | p104(X23)) & sP58(X23) & sP57(X23) & sP56(X23) & sP55(X23) & sP54(X23) & sP53(X23) & sP51(X23) & sP50(X23) & sP49(X23) & sP48(X23) & sP52(X23)) | ~ sP59(X23)), inference(nnf_transformation, [], [f67])).
fof(f67, plain, ! [X23] : (((~ p101(X23) | p100(X23)) & (~ p102(X23) | p101(X23)) & (~ p103(X23) | p102(X23)) & (~ p104(X23) | p103(X23)) & (~ p105(X23) | p104(X23)) & sP58(X23) & sP57(X23) & sP56(X23) & sP55(X23) & sP54(X23) & sP53(X23) & sP51(X23) & sP50(X23) & sP49(X23) & sP48(X23) & sP52(X23)) | ~ sP59(X23)), inference(usedef, [], [e67])).
fof(e67, plain, ! [X23] : (sP59(X23) <=> ((~ p101(X23) | p100(X23)) & (~ p102(X23) | p101(X23)) & (~ p103(X23) | p102(X23)) & (~ p104(X23) | p103(X23)) & (~ p105(X23) | p104(X23)) & sP58(X23) & sP57(X23) & sP56(X23) & sP55(X23) & sP54(X23) & sP53(X23) & sP51(X23) & sP50(X23) & sP49(X23) & sP48(X23) & sP52(X23))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f704, plain, sP59(sK111), inference(resolution, [], [f612, f695])).
fof(f695, plain, r1(sK110, sK111), inference(subsumption_resolution, [], [f694, f668])).
fof(f668, plain, ~ p101(sK110), inference(cnf_transformation, [], [f276])).
fof(f276, plain, (p100(sK110) & ~ p101(sK110) & (~ p101(sK110) | p100(sK110)) & (~ p102(sK110) | p101(sK110)) & (~ p103(sK110) | p102(sK110)) & (~ p104(sK110) | p103(sK110)) & (~ p105(sK110) | p104(sK110)) & (~ p100(sK110) | ((~ p1(sK110) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(sK110, X1))) & (p1(sK110) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(sK110, X2))))) & (~ p101(sK110) | ((~ p2(sK110) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(sK110, X3))) & (p2(sK110) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(sK110, X4))))) & (~ p102(sK110) | ((~ p3(sK110) | ! [X5] : (~ p102(X5) | p3(X5) | ~ r1(sK110, X5))) & (p3(sK110) | ! [X6] : (~ p102(X6) | ~ p3(X6) | ~ r1(sK110, X6))))) & (~ p103(sK110) | ((~ p4(sK110) | ! [X7] : (~ p103(X7) | p4(X7) | ~ r1(sK110, X7))) & (p4(sK110) | ! [X8] : (~ p103(X8) | ~ p4(X8) | ~ r1(sK110, X8))))) & (~ p104(sK110) | ((~ p5(sK110) | ! [X9] : (~ p104(X9) | p5(X9) | ~ r1(sK110, X9))) & (p5(sK110) | ! [X10] : (~ p104(X10) | ~ p5(X10) | ~ r1(sK110, X10))))) & (~ p105(sK110) | ((~ p6(sK110) | ! [X11] : (~ p105(X11) | p6(X11) | ~ r1(sK110, X11))) & (p6(sK110) | ! [X12] : (~ p105(X12) | ~ p6(X12) | ~ r1(sK110, X12))))) & (~ p100(sK110) | p101(sK110) | ((p101(sK111) & ~ p102(sK111) & p2(sK111) & r1(sK110, sK111)) & (p101(sK112) & ~ p102(sK112) & ~ p2(sK112) & r1(sK110, sK112)))) & (~ p101(sK110) | p102(sK110) | ((p102(sK113) & ~ p103(sK113) & p3(sK113) & r1(sK110, sK113)) & (p102(sK114) & ~ p103(sK114) & ~ p3(sK114) & r1(sK110, sK114)))) & (~ p102(sK110) | p103(sK110) | ((p103(sK115) & ~ p104(sK115) & p4(sK115) & r1(sK110, sK115)) & (p103(sK116) & ~ p104(sK116) & ~ p4(sK116) & r1(sK110, sK116)))) & (~ p103(sK110) | p104(sK110) | ((p104(sK117) & ~ p105(sK117) & p5(sK117) & r1(sK110, sK117)) & (p104(sK118) & ~ p105(sK118) & ~ p5(sK118) & r1(sK110, sK118)))) & (~ p104(sK110) | p105(sK110) | ((p105(sK119) & p6(sK119) & r1(sK110, sK119)) & (p105(sK120) & ~ p6(sK120) & r1(sK110, sK120)))) & ! [X23] : (sP59(X23) | ~ r1(sK110, X23)) & ! [X24] : (! [X25] : (sP47(X25) | ~ r1(X24, X25)) | ~ r1(sK110, X24)) & ! [X26] : (! [X27] : (! [X28] : (sP35(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) | ~ r1(sK110, X26)) & ! [X29] : (! [X30] : (! [X31] : (! [X32] : (sP23(X32) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(sK110, X29)) & ! [X33] : (! [X34] : (! [X35] : (! [X36] : (! [X37] : (sP11(X37) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X33, X34)) | ~ r1(sK110, X33)) & ! [X38] : (! [X39] : (! [X40] : (! [X41] : (! [X42] : (p3(X42) | ~ r1(X41, X42)) | ~ r1(X40, X41)) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(sK110, X38))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK110, sK111, sK112, sK113, sK114, sK115, sK116, sK117, sK118, sK119, sK120])], [f264, f275, f274, f273, f272, f271, f270, f269, f268, f267, f266, f265])).
fof(f265, plain, (? [X0] : (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p102(X0) | ((~ p3(X0) | ! [X5] : (~ p102(X5) | p3(X5) | ~ r1(X0, X5))) & (p3(X0) | ! [X6] : (~ p102(X6) | ~ p3(X6) | ~ r1(X0, X6))))) & (~ p103(X0) | ((~ p4(X0) | ! [X7] : (~ p103(X7) | p4(X7) | ~ r1(X0, X7))) & (p4(X0) | ! [X8] : (~ p103(X8) | ~ p4(X8) | ~ r1(X0, X8))))) & (~ p104(X0) | ((~ p5(X0) | ! [X9] : (~ p104(X9) | p5(X9) | ~ r1(X0, X9))) & (p5(X0) | ! [X10] : (~ p104(X10) | ~ p5(X10) | ~ r1(X0, X10))))) & (~ p105(X0) | ((~ p6(X0) | ! [X11] : (~ p105(X11) | p6(X11) | ~ r1(X0, X11))) & (p6(X0) | ! [X12] : (~ p105(X12) | ~ p6(X12) | ~ r1(X0, X12))))) & (~ p100(X0) | p101(X0) | (? [X13] : (p101(X13) & ~ p102(X13) & p2(X13) & r1(X0, X13)) & ? [X14] : (p101(X14) & ~ p102(X14) & ~ p2(X14) & r1(X0, X14)))) & (~ p101(X0) | p102(X0) | (? [X15] : (p102(X15) & ~ p103(X15) & p3(X15) & r1(X0, X15)) & ? [X16] : (p102(X16) & ~ p103(X16) & ~ p3(X16) & r1(X0, X16)))) & (~ p102(X0) | p103(X0) | (? [X17] : (p103(X17) & ~ p104(X17) & p4(X17) & r1(X0, X17)) & ? [X18] : (p103(X18) & ~ p104(X18) & ~ p4(X18) & r1(X0, X18)))) & (~ p103(X0) | p104(X0) | (? [X19] : (p104(X19) & ~ p105(X19) & p5(X19) & r1(X0, X19)) & ? [X20] : (p104(X20) & ~ p105(X20) & ~ p5(X20) & r1(X0, X20)))) & (~ p104(X0) | p105(X0) | (? [X21] : (p105(X21) & p6(X21) & r1(X0, X21)) & ? [X22] : (p105(X22) & ~ p6(X22) & r1(X0, X22)))) & ! [X23] : (sP59(X23) | ~ r1(X0, X23)) & ! [X24] : (! [X25] : (sP47(X25) | ~ r1(X24, X25)) | ~ r1(X0, X24)) & ! [X26] : (! [X27] : (! [X28] : (sP35(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) | ~ r1(X0, X26)) & ! [X29] : (! [X30] : (! [X31] : (! [X32] : (sP23(X32) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X0, X29)) & ! [X33] : (! [X34] : (! [X35] : (! [X36] : (! [X37] : (sP11(X37) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X33, X34)) | ~ r1(X0, X33)) & ! [X38] : (! [X39] : (! [X40] : (! [X41] : (! [X42] : (p3(X42) | ~ r1(X41, X42)) | ~ r1(X40, X41)) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(X0, X38))) => (p100(sK110) & ~ p101(sK110) & (~ p101(sK110) | p100(sK110)) & (~ p102(sK110) | p101(sK110)) & (~ p103(sK110) | p102(sK110)) & (~ p104(sK110) | p103(sK110)) & (~ p105(sK110) | p104(sK110)) & (~ p100(sK110) | ((~ p1(sK110) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(sK110, X1))) & (p1(sK110) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(sK110, X2))))) & (~ p101(sK110) | ((~ p2(sK110) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(sK110, X3))) & (p2(sK110) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(sK110, X4))))) & (~ p102(sK110) | ((~ p3(sK110) | ! [X5] : (~ p102(X5) | p3(X5) | ~ r1(sK110, X5))) & (p3(sK110) | ! [X6] : (~ p102(X6) | ~ p3(X6) | ~ r1(sK110, X6))))) & (~ p103(sK110) | ((~ p4(sK110) | ! [X7] : (~ p103(X7) | p4(X7) | ~ r1(sK110, X7))) & (p4(sK110) | ! [X8] : (~ p103(X8) | ~ p4(X8) | ~ r1(sK110, X8))))) & (~ p104(sK110) | ((~ p5(sK110) | ! [X9] : (~ p104(X9) | p5(X9) | ~ r1(sK110, X9))) & (p5(sK110) | ! [X10] : (~ p104(X10) | ~ p5(X10) | ~ r1(sK110, X10))))) & (~ p105(sK110) | ((~ p6(sK110) | ! [X11] : (~ p105(X11) | p6(X11) | ~ r1(sK110, X11))) & (p6(sK110) | ! [X12] : (~ p105(X12) | ~ p6(X12) | ~ r1(sK110, X12))))) & (~ p100(sK110) | p101(sK110) | (? [X13] : (p101(X13) & ~ p102(X13) & p2(X13) & r1(sK110, X13)) & ? [X14] : (p101(X14) & ~ p102(X14) & ~ p2(X14) & r1(sK110, X14)))) & (~ p101(sK110) | p102(sK110) | (? [X15] : (p102(X15) & ~ p103(X15) & p3(X15) & r1(sK110, X15)) & ? [X16] : (p102(X16) & ~ p103(X16) & ~ p3(X16) & r1(sK110, X16)))) & (~ p102(sK110) | p103(sK110) | (? [X17] : (p103(X17) & ~ p104(X17) & p4(X17) & r1(sK110, X17)) & ? [X18] : (p103(X18) & ~ p104(X18) & ~ p4(X18) & r1(sK110, X18)))) & (~ p103(sK110) | p104(sK110) | (? [X19] : (p104(X19) & ~ p105(X19) & p5(X19) & r1(sK110, X19)) & ? [X20] : (p104(X20) & ~ p105(X20) & ~ p5(X20) & r1(sK110, X20)))) & (~ p104(sK110) | p105(sK110) | (? [X21] : (p105(X21) & p6(X21) & r1(sK110, X21)) & ? [X22] : (p105(X22) & ~ p6(X22) & r1(sK110, X22)))) & ! [X23] : (sP59(X23) | ~ r1(sK110, X23)) & ! [X24] : (! [X25] : (sP47(X25) | ~ r1(X24, X25)) | ~ r1(sK110, X24)) & ! [X26] : (! [X27] : (! [X28] : (sP35(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) | ~ r1(sK110, X26)) & ! [X29] : (! [X30] : (! [X31] : (! [X32] : (sP23(X32) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(sK110, X29)) & ! [X33] : (! [X34] : (! [X35] : (! [X36] : (! [X37] : (sP11(X37) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X33, X34)) | ~ r1(sK110, X33)) & ! [X38] : (! [X39] : (! [X40] : (! [X41] : (! [X42] : (p3(X42) | ~ r1(X41, X42)) | ~ r1(X40, X41)) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(sK110, X38)))), introduced(choice_axiom, [])).
fof(f266, plain, (? [X13] : (p101(X13) & ~ p102(X13) & p2(X13) & r1(sK110, X13)) => (p101(sK111) & ~ p102(sK111) & p2(sK111) & r1(sK110, sK111))), introduced(choice_axiom, [])).
fof(f267, plain, (? [X14] : (p101(X14) & ~ p102(X14) & ~ p2(X14) & r1(sK110, X14)) => (p101(sK112) & ~ p102(sK112) & ~ p2(sK112) & r1(sK110, sK112))), introduced(choice_axiom, [])).
fof(f268, plain, (? [X15] : (p102(X15) & ~ p103(X15) & p3(X15) & r1(sK110, X15)) => (p102(sK113) & ~ p103(sK113) & p3(sK113) & r1(sK110, sK113))), introduced(choice_axiom, [])).
fof(f269, plain, (? [X16] : (p102(X16) & ~ p103(X16) & ~ p3(X16) & r1(sK110, X16)) => (p102(sK114) & ~ p103(sK114) & ~ p3(sK114) & r1(sK110, sK114))), introduced(choice_axiom, [])).
fof(f270, plain, (? [X17] : (p103(X17) & ~ p104(X17) & p4(X17) & r1(sK110, X17)) => (p103(sK115) & ~ p104(sK115) & p4(sK115) & r1(sK110, sK115))), introduced(choice_axiom, [])).
fof(f271, plain, (? [X18] : (p103(X18) & ~ p104(X18) & ~ p4(X18) & r1(sK110, X18)) => (p103(sK116) & ~ p104(sK116) & ~ p4(sK116) & r1(sK110, sK116))), introduced(choice_axiom, [])).
fof(f272, plain, (? [X19] : (p104(X19) & ~ p105(X19) & p5(X19) & r1(sK110, X19)) => (p104(sK117) & ~ p105(sK117) & p5(sK117) & r1(sK110, sK117))), introduced(choice_axiom, [])).
fof(f273, plain, (? [X20] : (p104(X20) & ~ p105(X20) & ~ p5(X20) & r1(sK110, X20)) => (p104(sK118) & ~ p105(sK118) & ~ p5(sK118) & r1(sK110, sK118))), introduced(choice_axiom, [])).
fof(f274, plain, (? [X21] : (p105(X21) & p6(X21) & r1(sK110, X21)) => (p105(sK119) & p6(sK119) & r1(sK110, sK119))), introduced(choice_axiom, [])).
fof(f275, plain, (? [X22] : (p105(X22) & ~ p6(X22) & r1(sK110, X22)) => (p105(sK120) & ~ p6(sK120) & r1(sK110, sK120))), introduced(choice_axiom, [])).
fof(f264, plain, ? [X0] : (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p102(X0) | ((~ p3(X0) | ! [X5] : (~ p102(X5) | p3(X5) | ~ r1(X0, X5))) & (p3(X0) | ! [X6] : (~ p102(X6) | ~ p3(X6) | ~ r1(X0, X6))))) & (~ p103(X0) | ((~ p4(X0) | ! [X7] : (~ p103(X7) | p4(X7) | ~ r1(X0, X7))) & (p4(X0) | ! [X8] : (~ p103(X8) | ~ p4(X8) | ~ r1(X0, X8))))) & (~ p104(X0) | ((~ p5(X0) | ! [X9] : (~ p104(X9) | p5(X9) | ~ r1(X0, X9))) & (p5(X0) | ! [X10] : (~ p104(X10) | ~ p5(X10) | ~ r1(X0, X10))))) & (~ p105(X0) | ((~ p6(X0) | ! [X11] : (~ p105(X11) | p6(X11) | ~ r1(X0, X11))) & (p6(X0) | ! [X12] : (~ p105(X12) | ~ p6(X12) | ~ r1(X0, X12))))) & (~ p100(X0) | p101(X0) | (? [X13] : (p101(X13) & ~ p102(X13) & p2(X13) & r1(X0, X13)) & ? [X14] : (p101(X14) & ~ p102(X14) & ~ p2(X14) & r1(X0, X14)))) & (~ p101(X0) | p102(X0) | (? [X15] : (p102(X15) & ~ p103(X15) & p3(X15) & r1(X0, X15)) & ? [X16] : (p102(X16) & ~ p103(X16) & ~ p3(X16) & r1(X0, X16)))) & (~ p102(X0) | p103(X0) | (? [X17] : (p103(X17) & ~ p104(X17) & p4(X17) & r1(X0, X17)) & ? [X18] : (p103(X18) & ~ p104(X18) & ~ p4(X18) & r1(X0, X18)))) & (~ p103(X0) | p104(X0) | (? [X19] : (p104(X19) & ~ p105(X19) & p5(X19) & r1(X0, X19)) & ? [X20] : (p104(X20) & ~ p105(X20) & ~ p5(X20) & r1(X0, X20)))) & (~ p104(X0) | p105(X0) | (? [X21] : (p105(X21) & p6(X21) & r1(X0, X21)) & ? [X22] : (p105(X22) & ~ p6(X22) & r1(X0, X22)))) & ! [X23] : (sP59(X23) | ~ r1(X0, X23)) & ! [X24] : (! [X25] : (sP47(X25) | ~ r1(X24, X25)) | ~ r1(X0, X24)) & ! [X26] : (! [X27] : (! [X28] : (sP35(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) | ~ r1(X0, X26)) & ! [X29] : (! [X30] : (! [X31] : (! [X32] : (sP23(X32) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X0, X29)) & ! [X33] : (! [X34] : (! [X35] : (! [X36] : (! [X37] : (sP11(X37) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X33, X34)) | ~ r1(X0, X33)) & ! [X38] : (! [X39] : (! [X40] : (! [X41] : (! [X42] : (p3(X42) | ~ r1(X41, X42)) | ~ r1(X40, X41)) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(X0, X38))), inference(rectify, [], [f68])).
fof(f68, plain, ? [X0] : (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p102(X0) | ((~ p3(X0) | ! [X5] : (~ p102(X5) | p3(X5) | ~ r1(X0, X5))) & (p3(X0) | ! [X6] : (~ p102(X6) | ~ p3(X6) | ~ r1(X0, X6))))) & (~ p103(X0) | ((~ p4(X0) | ! [X7] : (~ p103(X7) | p4(X7) | ~ r1(X0, X7))) & (p4(X0) | ! [X8] : (~ p103(X8) | ~ p4(X8) | ~ r1(X0, X8))))) & (~ p104(X0) | ((~ p5(X0) | ! [X9] : (~ p104(X9) | p5(X9) | ~ r1(X0, X9))) & (p5(X0) | ! [X10] : (~ p104(X10) | ~ p5(X10) | ~ r1(X0, X10))))) & (~ p105(X0) | ((~ p6(X0) | ! [X11] : (~ p105(X11) | p6(X11) | ~ r1(X0, X11))) & (p6(X0) | ! [X12] : (~ p105(X12) | ~ p6(X12) | ~ r1(X0, X12))))) & (~ p100(X0) | p101(X0) | (? [X13] : (p101(X13) & ~ p102(X13) & p2(X13) & r1(X0, X13)) & ? [X14] : (p101(X14) & ~ p102(X14) & ~ p2(X14) & r1(X0, X14)))) & (~ p101(X0) | p102(X0) | (? [X15] : (p102(X15) & ~ p103(X15) & p3(X15) & r1(X0, X15)) & ? [X16] : (p102(X16) & ~ p103(X16) & ~ p3(X16) & r1(X0, X16)))) & (~ p102(X0) | p103(X0) | (? [X17] : (p103(X17) & ~ p104(X17) & p4(X17) & r1(X0, X17)) & ? [X18] : (p103(X18) & ~ p104(X18) & ~ p4(X18) & r1(X0, X18)))) & (~ p103(X0) | p104(X0) | (? [X19] : (p104(X19) & ~ p105(X19) & p5(X19) & r1(X0, X19)) & ? [X20] : (p104(X20) & ~ p105(X20) & ~ p5(X20) & r1(X0, X20)))) & (~ p104(X0) | p105(X0) | (? [X21] : (p105(X21) & p6(X21) & r1(X0, X21)) & ? [X22] : (p105(X22) & ~ p6(X22) & r1(X0, X22)))) & ! [X23] : (sP59(X23) | ~ r1(X0, X23)) & ! [X46] : (! [X47] : (sP47(X47) | ~ r1(X46, X47)) | ~ r1(X0, X46)) & ! [X70] : (! [X71] : (! [X72] : (sP35(X72) | ~ r1(X71, X72)) | ~ r1(X70, X71)) | ~ r1(X0, X70)) & ! [X95] : (! [X96] : (! [X97] : (! [X98] : (sP23(X98) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X0, X95)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (sP11(X125) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121)) & ! [X148] : (! [X149] : (! [X150] : (! [X151] : (! [X152] : (p3(X152) | ~ r1(X151, X152)) | ~ r1(X150, X151)) | ~ r1(X149, X150)) | ~ r1(X148, X149)) | ~ r1(X0, X148))), inference(definition_folding, [], [f7, e67, e66, e65, e64, e63, e62, e61, e60, e59, e58, e57, e56, e55, e54, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8])).
fof(f8, plain, ! [X125] : (~ p103(X125) | p104(X125) | (? [X144] : (p104(X144) & ~ p105(X144) & p5(X144) & r1(X125, X144)) & ? [X145] : (p104(X145) & ~ p105(X145) & ~ p5(X145) & r1(X125, X145))) | ~ sP0(X125)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X125] : (sP0(X125) <=> (~ p103(X125) | p104(X125) | (? [X144] : (p104(X144) & ~ p105(X144) & p5(X144) & r1(X125, X144)) & ? [X145] : (p104(X145) & ~ p105(X145) & ~ p5(X145) & r1(X125, X145))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f9, plain, ! [X125] : (~ p102(X125) | p103(X125) | (? [X142] : (p103(X142) & ~ p104(X142) & p4(X142) & r1(X125, X142)) & ? [X143] : (p103(X143) & ~ p104(X143) & ~ p4(X143) & r1(X125, X143))) | ~ sP1(X125)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X125] : (sP1(X125) <=> (~ p102(X125) | p103(X125) | (? [X142] : (p103(X142) & ~ p104(X142) & p4(X142) & r1(X125, X142)) & ? [X143] : (p103(X143) & ~ p104(X143) & ~ p4(X143) & r1(X125, X143))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f10, plain, ! [X125] : (~ p101(X125) | p102(X125) | (? [X140] : (p102(X140) & ~ p103(X140) & p3(X140) & r1(X125, X140)) & ? [X141] : (p102(X141) & ~ p103(X141) & ~ p3(X141) & r1(X125, X141))) | ~ sP2(X125)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X125] : (sP2(X125) <=> (~ p101(X125) | p102(X125) | (? [X140] : (p102(X140) & ~ p103(X140) & p3(X140) & r1(X125, X140)) & ? [X141] : (p102(X141) & ~ p103(X141) & ~ p3(X141) & r1(X125, X141))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f11, plain, ! [X125] : (~ p100(X125) | p101(X125) | (? [X138] : (p101(X138) & ~ p102(X138) & p2(X138) & r1(X125, X138)) & ? [X139] : (p101(X139) & ~ p102(X139) & ~ p2(X139) & r1(X125, X139))) | ~ sP3(X125)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X125] : (sP3(X125) <=> (~ p100(X125) | p101(X125) | (? [X138] : (p101(X138) & ~ p102(X138) & p2(X138) & r1(X125, X138)) & ? [X139] : (p101(X139) & ~ p102(X139) & ~ p2(X139) & r1(X125, X139))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f12, plain, ! [X125] : (~ p104(X125) | p105(X125) | (? [X146] : (p105(X146) & p6(X146) & r1(X125, X146)) & ? [X147] : (p105(X147) & ~ p6(X147) & r1(X125, X147))) | ~ sP4(X125)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X125] : (sP4(X125) <=> (~ p104(X125) | p105(X125) | (? [X146] : (p105(X146) & p6(X146) & r1(X125, X146)) & ? [X147] : (p105(X147) & ~ p6(X147) & r1(X125, X147))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f13, plain, ! [X125] : (~ p105(X125) | ((~ p6(X125) | ! [X136] : (~ p105(X136) | p6(X136) | ~ r1(X125, X136))) & (p6(X125) | ! [X137] : (~ p105(X137) | ~ p6(X137) | ~ r1(X125, X137)))) | ~ sP5(X125)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X125] : (sP5(X125) <=> (~ p105(X125) | ((~ p6(X125) | ! [X136] : (~ p105(X136) | p6(X136) | ~ r1(X125, X136))) & (p6(X125) | ! [X137] : (~ p105(X137) | ~ p6(X137) | ~ r1(X125, X137)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f14, plain, ! [X125] : (~ p104(X125) | ((~ p5(X125) | ! [X134] : (~ p104(X134) | p5(X134) | ~ r1(X125, X134))) & (p5(X125) | ! [X135] : (~ p104(X135) | ~ p5(X135) | ~ r1(X125, X135)))) | ~ sP6(X125)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X125] : (sP6(X125) <=> (~ p104(X125) | ((~ p5(X125) | ! [X134] : (~ p104(X134) | p5(X134) | ~ r1(X125, X134))) & (p5(X125) | ! [X135] : (~ p104(X135) | ~ p5(X135) | ~ r1(X125, X135)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f15, plain, ! [X125] : (~ p103(X125) | ((~ p4(X125) | ! [X132] : (~ p103(X132) | p4(X132) | ~ r1(X125, X132))) & (p4(X125) | ! [X133] : (~ p103(X133) | ~ p4(X133) | ~ r1(X125, X133)))) | ~ sP7(X125)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X125] : (sP7(X125) <=> (~ p103(X125) | ((~ p4(X125) | ! [X132] : (~ p103(X132) | p4(X132) | ~ r1(X125, X132))) & (p4(X125) | ! [X133] : (~ p103(X133) | ~ p4(X133) | ~ r1(X125, X133)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f16, plain, ! [X125] : (~ p102(X125) | ((~ p3(X125) | ! [X130] : (~ p102(X130) | p3(X130) | ~ r1(X125, X130))) & (p3(X125) | ! [X131] : (~ p102(X131) | ~ p3(X131) | ~ r1(X125, X131)))) | ~ sP8(X125)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X125] : (sP8(X125) <=> (~ p102(X125) | ((~ p3(X125) | ! [X130] : (~ p102(X130) | p3(X130) | ~ r1(X125, X130))) & (p3(X125) | ! [X131] : (~ p102(X131) | ~ p3(X131) | ~ r1(X125, X131)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f17, plain, ! [X125] : (~ p101(X125) | ((~ p2(X125) | ! [X128] : (~ p101(X128) | p2(X128) | ~ r1(X125, X128))) & (p2(X125) | ! [X129] : (~ p101(X129) | ~ p2(X129) | ~ r1(X125, X129)))) | ~ sP9(X125)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X125] : (sP9(X125) <=> (~ p101(X125) | ((~ p2(X125) | ! [X128] : (~ p101(X128) | p2(X128) | ~ r1(X125, X128))) & (p2(X125) | ! [X129] : (~ p101(X129) | ~ p2(X129) | ~ r1(X125, X129)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f18, plain, ! [X125] : (~ p100(X125) | ((~ p1(X125) | ! [X126] : (~ p100(X126) | p1(X126) | ~ r1(X125, X126))) & (p1(X125) | ! [X127] : (~ p100(X127) | ~ p1(X127) | ~ r1(X125, X127)))) | ~ sP10(X125)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X125] : (sP10(X125) <=> (~ p100(X125) | ((~ p1(X125) | ! [X126] : (~ p100(X126) | p1(X126) | ~ r1(X125, X126))) & (p1(X125) | ! [X127] : (~ p100(X127) | ~ p1(X127) | ~ r1(X125, X127)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f19, plain, ! [X125] : (((~ p101(X125) | p100(X125)) & (~ p102(X125) | p101(X125)) & (~ p103(X125) | p102(X125)) & (~ p104(X125) | p103(X125)) & (~ p105(X125) | p104(X125)) & sP10(X125) & sP9(X125) & sP8(X125) & sP7(X125) & sP6(X125) & sP5(X125) & sP3(X125) & sP2(X125) & sP1(X125) & sP0(X125) & sP4(X125)) | ~ sP11(X125)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X125] : (sP11(X125) <=> ((~ p101(X125) | p100(X125)) & (~ p102(X125) | p101(X125)) & (~ p103(X125) | p102(X125)) & (~ p104(X125) | p103(X125)) & (~ p105(X125) | p104(X125)) & sP10(X125) & sP9(X125) & sP8(X125) & sP7(X125) & sP6(X125) & sP5(X125) & sP3(X125) & sP2(X125) & sP1(X125) & sP0(X125) & sP4(X125))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f20, plain, ! [X98] : (~ p103(X98) | p104(X98) | (? [X117] : (p104(X117) & ~ p105(X117) & p5(X117) & r1(X98, X117)) & ? [X118] : (p104(X118) & ~ p105(X118) & ~ p5(X118) & r1(X98, X118))) | ~ sP12(X98)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X98] : (sP12(X98) <=> (~ p103(X98) | p104(X98) | (? [X117] : (p104(X117) & ~ p105(X117) & p5(X117) & r1(X98, X117)) & ? [X118] : (p104(X118) & ~ p105(X118) & ~ p5(X118) & r1(X98, X118))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f21, plain, ! [X98] : (~ p102(X98) | p103(X98) | (? [X115] : (p103(X115) & ~ p104(X115) & p4(X115) & r1(X98, X115)) & ? [X116] : (p103(X116) & ~ p104(X116) & ~ p4(X116) & r1(X98, X116))) | ~ sP13(X98)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X98] : (sP13(X98) <=> (~ p102(X98) | p103(X98) | (? [X115] : (p103(X115) & ~ p104(X115) & p4(X115) & r1(X98, X115)) & ? [X116] : (p103(X116) & ~ p104(X116) & ~ p4(X116) & r1(X98, X116))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f22, plain, ! [X98] : (~ p101(X98) | p102(X98) | (? [X113] : (p102(X113) & ~ p103(X113) & p3(X113) & r1(X98, X113)) & ? [X114] : (p102(X114) & ~ p103(X114) & ~ p3(X114) & r1(X98, X114))) | ~ sP14(X98)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X98] : (sP14(X98) <=> (~ p101(X98) | p102(X98) | (? [X113] : (p102(X113) & ~ p103(X113) & p3(X113) & r1(X98, X113)) & ? [X114] : (p102(X114) & ~ p103(X114) & ~ p3(X114) & r1(X98, X114))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f23, plain, ! [X98] : (~ p100(X98) | p101(X98) | (? [X111] : (p101(X111) & ~ p102(X111) & p2(X111) & r1(X98, X111)) & ? [X112] : (p101(X112) & ~ p102(X112) & ~ p2(X112) & r1(X98, X112))) | ~ sP15(X98)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X98] : (sP15(X98) <=> (~ p100(X98) | p101(X98) | (? [X111] : (p101(X111) & ~ p102(X111) & p2(X111) & r1(X98, X111)) & ? [X112] : (p101(X112) & ~ p102(X112) & ~ p2(X112) & r1(X98, X112))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f24, plain, ! [X98] : (~ p104(X98) | p105(X98) | (? [X119] : (p105(X119) & p6(X119) & r1(X98, X119)) & ? [X120] : (p105(X120) & ~ p6(X120) & r1(X98, X120))) | ~ sP16(X98)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X98] : (sP16(X98) <=> (~ p104(X98) | p105(X98) | (? [X119] : (p105(X119) & p6(X119) & r1(X98, X119)) & ? [X120] : (p105(X120) & ~ p6(X120) & r1(X98, X120))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f25, plain, ! [X98] : (~ p105(X98) | ((~ p6(X98) | ! [X109] : (~ p105(X109) | p6(X109) | ~ r1(X98, X109))) & (p6(X98) | ! [X110] : (~ p105(X110) | ~ p6(X110) | ~ r1(X98, X110)))) | ~ sP17(X98)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X98] : (sP17(X98) <=> (~ p105(X98) | ((~ p6(X98) | ! [X109] : (~ p105(X109) | p6(X109) | ~ r1(X98, X109))) & (p6(X98) | ! [X110] : (~ p105(X110) | ~ p6(X110) | ~ r1(X98, X110)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f26, plain, ! [X98] : (~ p104(X98) | ((~ p5(X98) | ! [X107] : (~ p104(X107) | p5(X107) | ~ r1(X98, X107))) & (p5(X98) | ! [X108] : (~ p104(X108) | ~ p5(X108) | ~ r1(X98, X108)))) | ~ sP18(X98)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X98] : (sP18(X98) <=> (~ p104(X98) | ((~ p5(X98) | ! [X107] : (~ p104(X107) | p5(X107) | ~ r1(X98, X107))) & (p5(X98) | ! [X108] : (~ p104(X108) | ~ p5(X108) | ~ r1(X98, X108)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f27, plain, ! [X98] : (~ p103(X98) | ((~ p4(X98) | ! [X105] : (~ p103(X105) | p4(X105) | ~ r1(X98, X105))) & (p4(X98) | ! [X106] : (~ p103(X106) | ~ p4(X106) | ~ r1(X98, X106)))) | ~ sP19(X98)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X98] : (sP19(X98) <=> (~ p103(X98) | ((~ p4(X98) | ! [X105] : (~ p103(X105) | p4(X105) | ~ r1(X98, X105))) & (p4(X98) | ! [X106] : (~ p103(X106) | ~ p4(X106) | ~ r1(X98, X106)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f28, plain, ! [X98] : (~ p102(X98) | ((~ p3(X98) | ! [X103] : (~ p102(X103) | p3(X103) | ~ r1(X98, X103))) & (p3(X98) | ! [X104] : (~ p102(X104) | ~ p3(X104) | ~ r1(X98, X104)))) | ~ sP20(X98)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X98] : (sP20(X98) <=> (~ p102(X98) | ((~ p3(X98) | ! [X103] : (~ p102(X103) | p3(X103) | ~ r1(X98, X103))) & (p3(X98) | ! [X104] : (~ p102(X104) | ~ p3(X104) | ~ r1(X98, X104)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f29, plain, ! [X98] : (~ p101(X98) | ((~ p2(X98) | ! [X101] : (~ p101(X101) | p2(X101) | ~ r1(X98, X101))) & (p2(X98) | ! [X102] : (~ p101(X102) | ~ p2(X102) | ~ r1(X98, X102)))) | ~ sP21(X98)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X98] : (sP21(X98) <=> (~ p101(X98) | ((~ p2(X98) | ! [X101] : (~ p101(X101) | p2(X101) | ~ r1(X98, X101))) & (p2(X98) | ! [X102] : (~ p101(X102) | ~ p2(X102) | ~ r1(X98, X102)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f30, plain, ! [X98] : (~ p100(X98) | ((~ p1(X98) | ! [X99] : (~ p100(X99) | p1(X99) | ~ r1(X98, X99))) & (p1(X98) | ! [X100] : (~ p100(X100) | ~ p1(X100) | ~ r1(X98, X100)))) | ~ sP22(X98)), inference(usedef, [], [e30])).
fof(e30, plain, ! [X98] : (sP22(X98) <=> (~ p100(X98) | ((~ p1(X98) | ! [X99] : (~ p100(X99) | p1(X99) | ~ r1(X98, X99))) & (p1(X98) | ! [X100] : (~ p100(X100) | ~ p1(X100) | ~ r1(X98, X100)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f31, plain, ! [X98] : (((~ p101(X98) | p100(X98)) & (~ p102(X98) | p101(X98)) & (~ p103(X98) | p102(X98)) & (~ p104(X98) | p103(X98)) & (~ p105(X98) | p104(X98)) & sP22(X98) & sP21(X98) & sP20(X98) & sP19(X98) & sP18(X98) & sP17(X98) & sP15(X98) & sP14(X98) & sP13(X98) & sP12(X98) & sP16(X98)) | ~ sP23(X98)), inference(usedef, [], [e31])).
fof(e31, plain, ! [X98] : (sP23(X98) <=> ((~ p101(X98) | p100(X98)) & (~ p102(X98) | p101(X98)) & (~ p103(X98) | p102(X98)) & (~ p104(X98) | p103(X98)) & (~ p105(X98) | p104(X98)) & sP22(X98) & sP21(X98) & sP20(X98) & sP19(X98) & sP18(X98) & sP17(X98) & sP15(X98) & sP14(X98) & sP13(X98) & sP12(X98) & sP16(X98))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f32, plain, ! [X72] : (~ p103(X72) | p104(X72) | (? [X91] : (p104(X91) & ~ p105(X91) & p5(X91) & r1(X72, X91)) & ? [X92] : (p104(X92) & ~ p105(X92) & ~ p5(X92) & r1(X72, X92))) | ~ sP24(X72)), inference(usedef, [], [e32])).
fof(e32, plain, ! [X72] : (sP24(X72) <=> (~ p103(X72) | p104(X72) | (? [X91] : (p104(X91) & ~ p105(X91) & p5(X91) & r1(X72, X91)) & ? [X92] : (p104(X92) & ~ p105(X92) & ~ p5(X92) & r1(X72, X92))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f33, plain, ! [X72] : (~ p102(X72) | p103(X72) | (? [X89] : (p103(X89) & ~ p104(X89) & p4(X89) & r1(X72, X89)) & ? [X90] : (p103(X90) & ~ p104(X90) & ~ p4(X90) & r1(X72, X90))) | ~ sP25(X72)), inference(usedef, [], [e33])).
fof(e33, plain, ! [X72] : (sP25(X72) <=> (~ p102(X72) | p103(X72) | (? [X89] : (p103(X89) & ~ p104(X89) & p4(X89) & r1(X72, X89)) & ? [X90] : (p103(X90) & ~ p104(X90) & ~ p4(X90) & r1(X72, X90))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f34, plain, ! [X72] : (~ p101(X72) | p102(X72) | (? [X87] : (p102(X87) & ~ p103(X87) & p3(X87) & r1(X72, X87)) & ? [X88] : (p102(X88) & ~ p103(X88) & ~ p3(X88) & r1(X72, X88))) | ~ sP26(X72)), inference(usedef, [], [e34])).
fof(e34, plain, ! [X72] : (sP26(X72) <=> (~ p101(X72) | p102(X72) | (? [X87] : (p102(X87) & ~ p103(X87) & p3(X87) & r1(X72, X87)) & ? [X88] : (p102(X88) & ~ p103(X88) & ~ p3(X88) & r1(X72, X88))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f35, plain, ! [X72] : (~ p100(X72) | p101(X72) | (? [X85] : (p101(X85) & ~ p102(X85) & p2(X85) & r1(X72, X85)) & ? [X86] : (p101(X86) & ~ p102(X86) & ~ p2(X86) & r1(X72, X86))) | ~ sP27(X72)), inference(usedef, [], [e35])).
fof(e35, plain, ! [X72] : (sP27(X72) <=> (~ p100(X72) | p101(X72) | (? [X85] : (p101(X85) & ~ p102(X85) & p2(X85) & r1(X72, X85)) & ? [X86] : (p101(X86) & ~ p102(X86) & ~ p2(X86) & r1(X72, X86))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f36, plain, ! [X72] : (~ p104(X72) | p105(X72) | (? [X93] : (p105(X93) & p6(X93) & r1(X72, X93)) & ? [X94] : (p105(X94) & ~ p6(X94) & r1(X72, X94))) | ~ sP28(X72)), inference(usedef, [], [e36])).
fof(e36, plain, ! [X72] : (sP28(X72) <=> (~ p104(X72) | p105(X72) | (? [X93] : (p105(X93) & p6(X93) & r1(X72, X93)) & ? [X94] : (p105(X94) & ~ p6(X94) & r1(X72, X94))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f37, plain, ! [X72] : (~ p105(X72) | ((~ p6(X72) | ! [X83] : (~ p105(X83) | p6(X83) | ~ r1(X72, X83))) & (p6(X72) | ! [X84] : (~ p105(X84) | ~ p6(X84) | ~ r1(X72, X84)))) | ~ sP29(X72)), inference(usedef, [], [e37])).
fof(e37, plain, ! [X72] : (sP29(X72) <=> (~ p105(X72) | ((~ p6(X72) | ! [X83] : (~ p105(X83) | p6(X83) | ~ r1(X72, X83))) & (p6(X72) | ! [X84] : (~ p105(X84) | ~ p6(X84) | ~ r1(X72, X84)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f38, plain, ! [X72] : (~ p104(X72) | ((~ p5(X72) | ! [X81] : (~ p104(X81) | p5(X81) | ~ r1(X72, X81))) & (p5(X72) | ! [X82] : (~ p104(X82) | ~ p5(X82) | ~ r1(X72, X82)))) | ~ sP30(X72)), inference(usedef, [], [e38])).
fof(e38, plain, ! [X72] : (sP30(X72) <=> (~ p104(X72) | ((~ p5(X72) | ! [X81] : (~ p104(X81) | p5(X81) | ~ r1(X72, X81))) & (p5(X72) | ! [X82] : (~ p104(X82) | ~ p5(X82) | ~ r1(X72, X82)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f39, plain, ! [X72] : (~ p103(X72) | ((~ p4(X72) | ! [X79] : (~ p103(X79) | p4(X79) | ~ r1(X72, X79))) & (p4(X72) | ! [X80] : (~ p103(X80) | ~ p4(X80) | ~ r1(X72, X80)))) | ~ sP31(X72)), inference(usedef, [], [e39])).
fof(e39, plain, ! [X72] : (sP31(X72) <=> (~ p103(X72) | ((~ p4(X72) | ! [X79] : (~ p103(X79) | p4(X79) | ~ r1(X72, X79))) & (p4(X72) | ! [X80] : (~ p103(X80) | ~ p4(X80) | ~ r1(X72, X80)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f40, plain, ! [X72] : (~ p102(X72) | ((~ p3(X72) | ! [X77] : (~ p102(X77) | p3(X77) | ~ r1(X72, X77))) & (p3(X72) | ! [X78] : (~ p102(X78) | ~ p3(X78) | ~ r1(X72, X78)))) | ~ sP32(X72)), inference(usedef, [], [e40])).
fof(e40, plain, ! [X72] : (sP32(X72) <=> (~ p102(X72) | ((~ p3(X72) | ! [X77] : (~ p102(X77) | p3(X77) | ~ r1(X72, X77))) & (p3(X72) | ! [X78] : (~ p102(X78) | ~ p3(X78) | ~ r1(X72, X78)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f41, plain, ! [X72] : (~ p101(X72) | ((~ p2(X72) | ! [X75] : (~ p101(X75) | p2(X75) | ~ r1(X72, X75))) & (p2(X72) | ! [X76] : (~ p101(X76) | ~ p2(X76) | ~ r1(X72, X76)))) | ~ sP33(X72)), inference(usedef, [], [e41])).
fof(e41, plain, ! [X72] : (sP33(X72) <=> (~ p101(X72) | ((~ p2(X72) | ! [X75] : (~ p101(X75) | p2(X75) | ~ r1(X72, X75))) & (p2(X72) | ! [X76] : (~ p101(X76) | ~ p2(X76) | ~ r1(X72, X76)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f42, plain, ! [X72] : (~ p100(X72) | ((~ p1(X72) | ! [X73] : (~ p100(X73) | p1(X73) | ~ r1(X72, X73))) & (p1(X72) | ! [X74] : (~ p100(X74) | ~ p1(X74) | ~ r1(X72, X74)))) | ~ sP34(X72)), inference(usedef, [], [e42])).
fof(e42, plain, ! [X72] : (sP34(X72) <=> (~ p100(X72) | ((~ p1(X72) | ! [X73] : (~ p100(X73) | p1(X73) | ~ r1(X72, X73))) & (p1(X72) | ! [X74] : (~ p100(X74) | ~ p1(X74) | ~ r1(X72, X74)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f43, plain, ! [X72] : (((~ p101(X72) | p100(X72)) & (~ p102(X72) | p101(X72)) & (~ p103(X72) | p102(X72)) & (~ p104(X72) | p103(X72)) & (~ p105(X72) | p104(X72)) & sP34(X72) & sP33(X72) & sP32(X72) & sP31(X72) & sP30(X72) & sP29(X72) & sP27(X72) & sP26(X72) & sP25(X72) & sP24(X72) & sP28(X72)) | ~ sP35(X72)), inference(usedef, [], [e43])).
fof(e43, plain, ! [X72] : (sP35(X72) <=> ((~ p101(X72) | p100(X72)) & (~ p102(X72) | p101(X72)) & (~ p103(X72) | p102(X72)) & (~ p104(X72) | p103(X72)) & (~ p105(X72) | p104(X72)) & sP34(X72) & sP33(X72) & sP32(X72) & sP31(X72) & sP30(X72) & sP29(X72) & sP27(X72) & sP26(X72) & sP25(X72) & sP24(X72) & sP28(X72))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f44, plain, ! [X47] : (~ p103(X47) | p104(X47) | (? [X66] : (p104(X66) & ~ p105(X66) & p5(X66) & r1(X47, X66)) & ? [X67] : (p104(X67) & ~ p105(X67) & ~ p5(X67) & r1(X47, X67))) | ~ sP36(X47)), inference(usedef, [], [e44])).
fof(e44, plain, ! [X47] : (sP36(X47) <=> (~ p103(X47) | p104(X47) | (? [X66] : (p104(X66) & ~ p105(X66) & p5(X66) & r1(X47, X66)) & ? [X67] : (p104(X67) & ~ p105(X67) & ~ p5(X67) & r1(X47, X67))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f45, plain, ! [X47] : (~ p102(X47) | p103(X47) | (? [X64] : (p103(X64) & ~ p104(X64) & p4(X64) & r1(X47, X64)) & ? [X65] : (p103(X65) & ~ p104(X65) & ~ p4(X65) & r1(X47, X65))) | ~ sP37(X47)), inference(usedef, [], [e45])).
fof(e45, plain, ! [X47] : (sP37(X47) <=> (~ p102(X47) | p103(X47) | (? [X64] : (p103(X64) & ~ p104(X64) & p4(X64) & r1(X47, X64)) & ? [X65] : (p103(X65) & ~ p104(X65) & ~ p4(X65) & r1(X47, X65))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f46, plain, ! [X47] : (~ p101(X47) | p102(X47) | (? [X62] : (p102(X62) & ~ p103(X62) & p3(X62) & r1(X47, X62)) & ? [X63] : (p102(X63) & ~ p103(X63) & ~ p3(X63) & r1(X47, X63))) | ~ sP38(X47)), inference(usedef, [], [e46])).
fof(e46, plain, ! [X47] : (sP38(X47) <=> (~ p101(X47) | p102(X47) | (? [X62] : (p102(X62) & ~ p103(X62) & p3(X62) & r1(X47, X62)) & ? [X63] : (p102(X63) & ~ p103(X63) & ~ p3(X63) & r1(X47, X63))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f47, plain, ! [X47] : (~ p100(X47) | p101(X47) | (? [X60] : (p101(X60) & ~ p102(X60) & p2(X60) & r1(X47, X60)) & ? [X61] : (p101(X61) & ~ p102(X61) & ~ p2(X61) & r1(X47, X61))) | ~ sP39(X47)), inference(usedef, [], [e47])).
fof(e47, plain, ! [X47] : (sP39(X47) <=> (~ p100(X47) | p101(X47) | (? [X60] : (p101(X60) & ~ p102(X60) & p2(X60) & r1(X47, X60)) & ? [X61] : (p101(X61) & ~ p102(X61) & ~ p2(X61) & r1(X47, X61))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f48, plain, ! [X47] : (~ p104(X47) | p105(X47) | (? [X68] : (p105(X68) & p6(X68) & r1(X47, X68)) & ? [X69] : (p105(X69) & ~ p6(X69) & r1(X47, X69))) | ~ sP40(X47)), inference(usedef, [], [e48])).
fof(e48, plain, ! [X47] : (sP40(X47) <=> (~ p104(X47) | p105(X47) | (? [X68] : (p105(X68) & p6(X68) & r1(X47, X68)) & ? [X69] : (p105(X69) & ~ p6(X69) & r1(X47, X69))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f49, plain, ! [X47] : (~ p105(X47) | ((~ p6(X47) | ! [X58] : (~ p105(X58) | p6(X58) | ~ r1(X47, X58))) & (p6(X47) | ! [X59] : (~ p105(X59) | ~ p6(X59) | ~ r1(X47, X59)))) | ~ sP41(X47)), inference(usedef, [], [e49])).
fof(e49, plain, ! [X47] : (sP41(X47) <=> (~ p105(X47) | ((~ p6(X47) | ! [X58] : (~ p105(X58) | p6(X58) | ~ r1(X47, X58))) & (p6(X47) | ! [X59] : (~ p105(X59) | ~ p6(X59) | ~ r1(X47, X59)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f50, plain, ! [X47] : (~ p104(X47) | ((~ p5(X47) | ! [X56] : (~ p104(X56) | p5(X56) | ~ r1(X47, X56))) & (p5(X47) | ! [X57] : (~ p104(X57) | ~ p5(X57) | ~ r1(X47, X57)))) | ~ sP42(X47)), inference(usedef, [], [e50])).
fof(e50, plain, ! [X47] : (sP42(X47) <=> (~ p104(X47) | ((~ p5(X47) | ! [X56] : (~ p104(X56) | p5(X56) | ~ r1(X47, X56))) & (p5(X47) | ! [X57] : (~ p104(X57) | ~ p5(X57) | ~ r1(X47, X57)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f51, plain, ! [X47] : (~ p103(X47) | ((~ p4(X47) | ! [X54] : (~ p103(X54) | p4(X54) | ~ r1(X47, X54))) & (p4(X47) | ! [X55] : (~ p103(X55) | ~ p4(X55) | ~ r1(X47, X55)))) | ~ sP43(X47)), inference(usedef, [], [e51])).
fof(e51, plain, ! [X47] : (sP43(X47) <=> (~ p103(X47) | ((~ p4(X47) | ! [X54] : (~ p103(X54) | p4(X54) | ~ r1(X47, X54))) & (p4(X47) | ! [X55] : (~ p103(X55) | ~ p4(X55) | ~ r1(X47, X55)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f52, plain, ! [X47] : (~ p102(X47) | ((~ p3(X47) | ! [X52] : (~ p102(X52) | p3(X52) | ~ r1(X47, X52))) & (p3(X47) | ! [X53] : (~ p102(X53) | ~ p3(X53) | ~ r1(X47, X53)))) | ~ sP44(X47)), inference(usedef, [], [e52])).
fof(e52, plain, ! [X47] : (sP44(X47) <=> (~ p102(X47) | ((~ p3(X47) | ! [X52] : (~ p102(X52) | p3(X52) | ~ r1(X47, X52))) & (p3(X47) | ! [X53] : (~ p102(X53) | ~ p3(X53) | ~ r1(X47, X53)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f53, plain, ! [X47] : (~ p101(X47) | ((~ p2(X47) | ! [X50] : (~ p101(X50) | p2(X50) | ~ r1(X47, X50))) & (p2(X47) | ! [X51] : (~ p101(X51) | ~ p2(X51) | ~ r1(X47, X51)))) | ~ sP45(X47)), inference(usedef, [], [e53])).
fof(e53, plain, ! [X47] : (sP45(X47) <=> (~ p101(X47) | ((~ p2(X47) | ! [X50] : (~ p101(X50) | p2(X50) | ~ r1(X47, X50))) & (p2(X47) | ! [X51] : (~ p101(X51) | ~ p2(X51) | ~ r1(X47, X51)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f54, plain, ! [X47] : (~ p100(X47) | ((~ p1(X47) | ! [X48] : (~ p100(X48) | p1(X48) | ~ r1(X47, X48))) & (p1(X47) | ! [X49] : (~ p100(X49) | ~ p1(X49) | ~ r1(X47, X49)))) | ~ sP46(X47)), inference(usedef, [], [e54])).
fof(e54, plain, ! [X47] : (sP46(X47) <=> (~ p100(X47) | ((~ p1(X47) | ! [X48] : (~ p100(X48) | p1(X48) | ~ r1(X47, X48))) & (p1(X47) | ! [X49] : (~ p100(X49) | ~ p1(X49) | ~ r1(X47, X49)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f55, plain, ! [X47] : (((~ p101(X47) | p100(X47)) & (~ p102(X47) | p101(X47)) & (~ p103(X47) | p102(X47)) & (~ p104(X47) | p103(X47)) & (~ p105(X47) | p104(X47)) & sP46(X47) & sP45(X47) & sP44(X47) & sP43(X47) & sP42(X47) & sP41(X47) & sP39(X47) & sP38(X47) & sP37(X47) & sP36(X47) & sP40(X47)) | ~ sP47(X47)), inference(usedef, [], [e55])).
fof(e55, plain, ! [X47] : (sP47(X47) <=> ((~ p101(X47) | p100(X47)) & (~ p102(X47) | p101(X47)) & (~ p103(X47) | p102(X47)) & (~ p104(X47) | p103(X47)) & (~ p105(X47) | p104(X47)) & sP46(X47) & sP45(X47) & sP44(X47) & sP43(X47) & sP42(X47) & sP41(X47) & sP39(X47) & sP38(X47) & sP37(X47) & sP36(X47) & sP40(X47))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f56, plain, ! [X23] : (~ p103(X23) | p104(X23) | (? [X42] : (p104(X42) & ~ p105(X42) & p5(X42) & r1(X23, X42)) & ? [X43] : (p104(X43) & ~ p105(X43) & ~ p5(X43) & r1(X23, X43))) | ~ sP48(X23)), inference(usedef, [], [e56])).
fof(e56, plain, ! [X23] : (sP48(X23) <=> (~ p103(X23) | p104(X23) | (? [X42] : (p104(X42) & ~ p105(X42) & p5(X42) & r1(X23, X42)) & ? [X43] : (p104(X43) & ~ p105(X43) & ~ p5(X43) & r1(X23, X43))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f57, plain, ! [X23] : (~ p102(X23) | p103(X23) | (? [X40] : (p103(X40) & ~ p104(X40) & p4(X40) & r1(X23, X40)) & ? [X41] : (p103(X41) & ~ p104(X41) & ~ p4(X41) & r1(X23, X41))) | ~ sP49(X23)), inference(usedef, [], [e57])).
fof(e57, plain, ! [X23] : (sP49(X23) <=> (~ p102(X23) | p103(X23) | (? [X40] : (p103(X40) & ~ p104(X40) & p4(X40) & r1(X23, X40)) & ? [X41] : (p103(X41) & ~ p104(X41) & ~ p4(X41) & r1(X23, X41))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f58, plain, ! [X23] : (~ p101(X23) | p102(X23) | (? [X38] : (p102(X38) & ~ p103(X38) & p3(X38) & r1(X23, X38)) & ? [X39] : (p102(X39) & ~ p103(X39) & ~ p3(X39) & r1(X23, X39))) | ~ sP50(X23)), inference(usedef, [], [e58])).
fof(e58, plain, ! [X23] : (sP50(X23) <=> (~ p101(X23) | p102(X23) | (? [X38] : (p102(X38) & ~ p103(X38) & p3(X38) & r1(X23, X38)) & ? [X39] : (p102(X39) & ~ p103(X39) & ~ p3(X39) & r1(X23, X39))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f59, plain, ! [X23] : (~ p100(X23) | p101(X23) | (? [X36] : (p101(X36) & ~ p102(X36) & p2(X36) & r1(X23, X36)) & ? [X37] : (p101(X37) & ~ p102(X37) & ~ p2(X37) & r1(X23, X37))) | ~ sP51(X23)), inference(usedef, [], [e59])).
fof(e59, plain, ! [X23] : (sP51(X23) <=> (~ p100(X23) | p101(X23) | (? [X36] : (p101(X36) & ~ p102(X36) & p2(X36) & r1(X23, X36)) & ? [X37] : (p101(X37) & ~ p102(X37) & ~ p2(X37) & r1(X23, X37))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f60, plain, ! [X23] : (~ p104(X23) | p105(X23) | (? [X44] : (p105(X44) & p6(X44) & r1(X23, X44)) & ? [X45] : (p105(X45) & ~ p6(X45) & r1(X23, X45))) | ~ sP52(X23)), inference(usedef, [], [e60])).
fof(e60, plain, ! [X23] : (sP52(X23) <=> (~ p104(X23) | p105(X23) | (? [X44] : (p105(X44) & p6(X44) & r1(X23, X44)) & ? [X45] : (p105(X45) & ~ p6(X45) & r1(X23, X45))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f61, plain, ! [X23] : (~ p105(X23) | ((~ p6(X23) | ! [X34] : (~ p105(X34) | p6(X34) | ~ r1(X23, X34))) & (p6(X23) | ! [X35] : (~ p105(X35) | ~ p6(X35) | ~ r1(X23, X35)))) | ~ sP53(X23)), inference(usedef, [], [e61])).
fof(e61, plain, ! [X23] : (sP53(X23) <=> (~ p105(X23) | ((~ p6(X23) | ! [X34] : (~ p105(X34) | p6(X34) | ~ r1(X23, X34))) & (p6(X23) | ! [X35] : (~ p105(X35) | ~ p6(X35) | ~ r1(X23, X35)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f62, plain, ! [X23] : (~ p104(X23) | ((~ p5(X23) | ! [X32] : (~ p104(X32) | p5(X32) | ~ r1(X23, X32))) & (p5(X23) | ! [X33] : (~ p104(X33) | ~ p5(X33) | ~ r1(X23, X33)))) | ~ sP54(X23)), inference(usedef, [], [e62])).
fof(e62, plain, ! [X23] : (sP54(X23) <=> (~ p104(X23) | ((~ p5(X23) | ! [X32] : (~ p104(X32) | p5(X32) | ~ r1(X23, X32))) & (p5(X23) | ! [X33] : (~ p104(X33) | ~ p5(X33) | ~ r1(X23, X33)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f63, plain, ! [X23] : (~ p103(X23) | ((~ p4(X23) | ! [X30] : (~ p103(X30) | p4(X30) | ~ r1(X23, X30))) & (p4(X23) | ! [X31] : (~ p103(X31) | ~ p4(X31) | ~ r1(X23, X31)))) | ~ sP55(X23)), inference(usedef, [], [e63])).
fof(e63, plain, ! [X23] : (sP55(X23) <=> (~ p103(X23) | ((~ p4(X23) | ! [X30] : (~ p103(X30) | p4(X30) | ~ r1(X23, X30))) & (p4(X23) | ! [X31] : (~ p103(X31) | ~ p4(X31) | ~ r1(X23, X31)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f64, plain, ! [X23] : (~ p102(X23) | ((~ p3(X23) | ! [X28] : (~ p102(X28) | p3(X28) | ~ r1(X23, X28))) & (p3(X23) | ! [X29] : (~ p102(X29) | ~ p3(X29) | ~ r1(X23, X29)))) | ~ sP56(X23)), inference(usedef, [], [e64])).
fof(e64, plain, ! [X23] : (sP56(X23) <=> (~ p102(X23) | ((~ p3(X23) | ! [X28] : (~ p102(X28) | p3(X28) | ~ r1(X23, X28))) & (p3(X23) | ! [X29] : (~ p102(X29) | ~ p3(X29) | ~ r1(X23, X29)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f65, plain, ! [X23] : (~ p101(X23) | ((~ p2(X23) | ! [X26] : (~ p101(X26) | p2(X26) | ~ r1(X23, X26))) & (p2(X23) | ! [X27] : (~ p101(X27) | ~ p2(X27) | ~ r1(X23, X27)))) | ~ sP57(X23)), inference(usedef, [], [e65])).
fof(e65, plain, ! [X23] : (sP57(X23) <=> (~ p101(X23) | ((~ p2(X23) | ! [X26] : (~ p101(X26) | p2(X26) | ~ r1(X23, X26))) & (p2(X23) | ! [X27] : (~ p101(X27) | ~ p2(X27) | ~ r1(X23, X27)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f66, plain, ! [X23] : (~ p100(X23) | ((~ p1(X23) | ! [X24] : (~ p100(X24) | p1(X24) | ~ r1(X23, X24))) & (p1(X23) | ! [X25] : (~ p100(X25) | ~ p1(X25) | ~ r1(X23, X25)))) | ~ sP58(X23)), inference(usedef, [], [e66])).
fof(e66, plain, ! [X23] : (sP58(X23) <=> (~ p100(X23) | ((~ p1(X23) | ! [X24] : (~ p100(X24) | p1(X24) | ~ r1(X23, X24))) & (p1(X23) | ! [X25] : (~ p100(X25) | ~ p1(X25) | ~ r1(X23, X25)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f7, plain, ? [X0] : (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p102(X0) | ((~ p3(X0) | ! [X5] : (~ p102(X5) | p3(X5) | ~ r1(X0, X5))) & (p3(X0) | ! [X6] : (~ p102(X6) | ~ p3(X6) | ~ r1(X0, X6))))) & (~ p103(X0) | ((~ p4(X0) | ! [X7] : (~ p103(X7) | p4(X7) | ~ r1(X0, X7))) & (p4(X0) | ! [X8] : (~ p103(X8) | ~ p4(X8) | ~ r1(X0, X8))))) & (~ p104(X0) | ((~ p5(X0) | ! [X9] : (~ p104(X9) | p5(X9) | ~ r1(X0, X9))) & (p5(X0) | ! [X10] : (~ p104(X10) | ~ p5(X10) | ~ r1(X0, X10))))) & (~ p105(X0) | ((~ p6(X0) | ! [X11] : (~ p105(X11) | p6(X11) | ~ r1(X0, X11))) & (p6(X0) | ! [X12] : (~ p105(X12) | ~ p6(X12) | ~ r1(X0, X12))))) & (~ p100(X0) | p101(X0) | (? [X13] : (p101(X13) & ~ p102(X13) & p2(X13) & r1(X0, X13)) & ? [X14] : (p101(X14) & ~ p102(X14) & ~ p2(X14) & r1(X0, X14)))) & (~ p101(X0) | p102(X0) | (? [X15] : (p102(X15) & ~ p103(X15) & p3(X15) & r1(X0, X15)) & ? [X16] : (p102(X16) & ~ p103(X16) & ~ p3(X16) & r1(X0, X16)))) & (~ p102(X0) | p103(X0) | (? [X17] : (p103(X17) & ~ p104(X17) & p4(X17) & r1(X0, X17)) & ? [X18] : (p103(X18) & ~ p104(X18) & ~ p4(X18) & r1(X0, X18)))) & (~ p103(X0) | p104(X0) | (? [X19] : (p104(X19) & ~ p105(X19) & p5(X19) & r1(X0, X19)) & ? [X20] : (p104(X20) & ~ p105(X20) & ~ p5(X20) & r1(X0, X20)))) & (~ p104(X0) | p105(X0) | (? [X21] : (p105(X21) & p6(X21) & r1(X0, X21)) & ? [X22] : (p105(X22) & ~ p6(X22) & r1(X0, X22)))) & ! [X23] : (((~ p101(X23) | p100(X23)) & (~ p102(X23) | p101(X23)) & (~ p103(X23) | p102(X23)) & (~ p104(X23) | p103(X23)) & (~ p105(X23) | p104(X23)) & (~ p100(X23) | ((~ p1(X23) | ! [X24] : (~ p100(X24) | p1(X24) | ~ r1(X23, X24))) & (p1(X23) | ! [X25] : (~ p100(X25) | ~ p1(X25) | ~ r1(X23, X25))))) & (~ p101(X23) | ((~ p2(X23) | ! [X26] : (~ p101(X26) | p2(X26) | ~ r1(X23, X26))) & (p2(X23) | ! [X27] : (~ p101(X27) | ~ p2(X27) | ~ r1(X23, X27))))) & (~ p102(X23) | ((~ p3(X23) | ! [X28] : (~ p102(X28) | p3(X28) | ~ r1(X23, X28))) & (p3(X23) | ! [X29] : (~ p102(X29) | ~ p3(X29) | ~ r1(X23, X29))))) & (~ p103(X23) | ((~ p4(X23) | ! [X30] : (~ p103(X30) | p4(X30) | ~ r1(X23, X30))) & (p4(X23) | ! [X31] : (~ p103(X31) | ~ p4(X31) | ~ r1(X23, X31))))) & (~ p104(X23) | ((~ p5(X23) | ! [X32] : (~ p104(X32) | p5(X32) | ~ r1(X23, X32))) & (p5(X23) | ! [X33] : (~ p104(X33) | ~ p5(X33) | ~ r1(X23, X33))))) & (~ p105(X23) | ((~ p6(X23) | ! [X34] : (~ p105(X34) | p6(X34) | ~ r1(X23, X34))) & (p6(X23) | ! [X35] : (~ p105(X35) | ~ p6(X35) | ~ r1(X23, X35))))) & (~ p100(X23) | p101(X23) | (? [X36] : (p101(X36) & ~ p102(X36) & p2(X36) & r1(X23, X36)) & ? [X37] : (p101(X37) & ~ p102(X37) & ~ p2(X37) & r1(X23, X37)))) & (~ p101(X23) | p102(X23) | (? [X38] : (p102(X38) & ~ p103(X38) & p3(X38) & r1(X23, X38)) & ? [X39] : (p102(X39) & ~ p103(X39) & ~ p3(X39) & r1(X23, X39)))) & (~ p102(X23) | p103(X23) | (? [X40] : (p103(X40) & ~ p104(X40) & p4(X40) & r1(X23, X40)) & ? [X41] : (p103(X41) & ~ p104(X41) & ~ p4(X41) & r1(X23, X41)))) & (~ p103(X23) | p104(X23) | (? [X42] : (p104(X42) & ~ p105(X42) & p5(X42) & r1(X23, X42)) & ? [X43] : (p104(X43) & ~ p105(X43) & ~ p5(X43) & r1(X23, X43)))) & (~ p104(X23) | p105(X23) | (? [X44] : (p105(X44) & p6(X44) & r1(X23, X44)) & ? [X45] : (p105(X45) & ~ p6(X45) & r1(X23, X45))))) | ~ r1(X0, X23)) & ! [X46] : (! [X47] : (((~ p101(X47) | p100(X47)) & (~ p102(X47) | p101(X47)) & (~ p103(X47) | p102(X47)) & (~ p104(X47) | p103(X47)) & (~ p105(X47) | p104(X47)) & (~ p100(X47) | ((~ p1(X47) | ! [X48] : (~ p100(X48) | p1(X48) | ~ r1(X47, X48))) & (p1(X47) | ! [X49] : (~ p100(X49) | ~ p1(X49) | ~ r1(X47, X49))))) & (~ p101(X47) | ((~ p2(X47) | ! [X50] : (~ p101(X50) | p2(X50) | ~ r1(X47, X50))) & (p2(X47) | ! [X51] : (~ p101(X51) | ~ p2(X51) | ~ r1(X47, X51))))) & (~ p102(X47) | ((~ p3(X47) | ! [X52] : (~ p102(X52) | p3(X52) | ~ r1(X47, X52))) & (p3(X47) | ! [X53] : (~ p102(X53) | ~ p3(X53) | ~ r1(X47, X53))))) & (~ p103(X47) | ((~ p4(X47) | ! [X54] : (~ p103(X54) | p4(X54) | ~ r1(X47, X54))) & (p4(X47) | ! [X55] : (~ p103(X55) | ~ p4(X55) | ~ r1(X47, X55))))) & (~ p104(X47) | ((~ p5(X47) | ! [X56] : (~ p104(X56) | p5(X56) | ~ r1(X47, X56))) & (p5(X47) | ! [X57] : (~ p104(X57) | ~ p5(X57) | ~ r1(X47, X57))))) & (~ p105(X47) | ((~ p6(X47) | ! [X58] : (~ p105(X58) | p6(X58) | ~ r1(X47, X58))) & (p6(X47) | ! [X59] : (~ p105(X59) | ~ p6(X59) | ~ r1(X47, X59))))) & (~ p100(X47) | p101(X47) | (? [X60] : (p101(X60) & ~ p102(X60) & p2(X60) & r1(X47, X60)) & ? [X61] : (p101(X61) & ~ p102(X61) & ~ p2(X61) & r1(X47, X61)))) & (~ p101(X47) | p102(X47) | (? [X62] : (p102(X62) & ~ p103(X62) & p3(X62) & r1(X47, X62)) & ? [X63] : (p102(X63) & ~ p103(X63) & ~ p3(X63) & r1(X47, X63)))) & (~ p102(X47) | p103(X47) | (? [X64] : (p103(X64) & ~ p104(X64) & p4(X64) & r1(X47, X64)) & ? [X65] : (p103(X65) & ~ p104(X65) & ~ p4(X65) & r1(X47, X65)))) & (~ p103(X47) | p104(X47) | (? [X66] : (p104(X66) & ~ p105(X66) & p5(X66) & r1(X47, X66)) & ? [X67] : (p104(X67) & ~ p105(X67) & ~ p5(X67) & r1(X47, X67)))) & (~ p104(X47) | p105(X47) | (? [X68] : (p105(X68) & p6(X68) & r1(X47, X68)) & ? [X69] : (p105(X69) & ~ p6(X69) & r1(X47, X69))))) | ~ r1(X46, X47)) | ~ r1(X0, X46)) & ! [X70] : (! [X71] : (! [X72] : (((~ p101(X72) | p100(X72)) & (~ p102(X72) | p101(X72)) & (~ p103(X72) | p102(X72)) & (~ p104(X72) | p103(X72)) & (~ p105(X72) | p104(X72)) & (~ p100(X72) | ((~ p1(X72) | ! [X73] : (~ p100(X73) | p1(X73) | ~ r1(X72, X73))) & (p1(X72) | ! [X74] : (~ p100(X74) | ~ p1(X74) | ~ r1(X72, X74))))) & (~ p101(X72) | ((~ p2(X72) | ! [X75] : (~ p101(X75) | p2(X75) | ~ r1(X72, X75))) & (p2(X72) | ! [X76] : (~ p101(X76) | ~ p2(X76) | ~ r1(X72, X76))))) & (~ p102(X72) | ((~ p3(X72) | ! [X77] : (~ p102(X77) | p3(X77) | ~ r1(X72, X77))) & (p3(X72) | ! [X78] : (~ p102(X78) | ~ p3(X78) | ~ r1(X72, X78))))) & (~ p103(X72) | ((~ p4(X72) | ! [X79] : (~ p103(X79) | p4(X79) | ~ r1(X72, X79))) & (p4(X72) | ! [X80] : (~ p103(X80) | ~ p4(X80) | ~ r1(X72, X80))))) & (~ p104(X72) | ((~ p5(X72) | ! [X81] : (~ p104(X81) | p5(X81) | ~ r1(X72, X81))) & (p5(X72) | ! [X82] : (~ p104(X82) | ~ p5(X82) | ~ r1(X72, X82))))) & (~ p105(X72) | ((~ p6(X72) | ! [X83] : (~ p105(X83) | p6(X83) | ~ r1(X72, X83))) & (p6(X72) | ! [X84] : (~ p105(X84) | ~ p6(X84) | ~ r1(X72, X84))))) & (~ p100(X72) | p101(X72) | (? [X85] : (p101(X85) & ~ p102(X85) & p2(X85) & r1(X72, X85)) & ? [X86] : (p101(X86) & ~ p102(X86) & ~ p2(X86) & r1(X72, X86)))) & (~ p101(X72) | p102(X72) | (? [X87] : (p102(X87) & ~ p103(X87) & p3(X87) & r1(X72, X87)) & ? [X88] : (p102(X88) & ~ p103(X88) & ~ p3(X88) & r1(X72, X88)))) & (~ p102(X72) | p103(X72) | (? [X89] : (p103(X89) & ~ p104(X89) & p4(X89) & r1(X72, X89)) & ? [X90] : (p103(X90) & ~ p104(X90) & ~ p4(X90) & r1(X72, X90)))) & (~ p103(X72) | p104(X72) | (? [X91] : (p104(X91) & ~ p105(X91) & p5(X91) & r1(X72, X91)) & ? [X92] : (p104(X92) & ~ p105(X92) & ~ p5(X92) & r1(X72, X92)))) & (~ p104(X72) | p105(X72) | (? [X93] : (p105(X93) & p6(X93) & r1(X72, X93)) & ? [X94] : (p105(X94) & ~ p6(X94) & r1(X72, X94))))) | ~ r1(X71, X72)) | ~ r1(X70, X71)) | ~ r1(X0, X70)) & ! [X95] : (! [X96] : (! [X97] : (! [X98] : (((~ p101(X98) | p100(X98)) & (~ p102(X98) | p101(X98)) & (~ p103(X98) | p102(X98)) & (~ p104(X98) | p103(X98)) & (~ p105(X98) | p104(X98)) & (~ p100(X98) | ((~ p1(X98) | ! [X99] : (~ p100(X99) | p1(X99) | ~ r1(X98, X99))) & (p1(X98) | ! [X100] : (~ p100(X100) | ~ p1(X100) | ~ r1(X98, X100))))) & (~ p101(X98) | ((~ p2(X98) | ! [X101] : (~ p101(X101) | p2(X101) | ~ r1(X98, X101))) & (p2(X98) | ! [X102] : (~ p101(X102) | ~ p2(X102) | ~ r1(X98, X102))))) & (~ p102(X98) | ((~ p3(X98) | ! [X103] : (~ p102(X103) | p3(X103) | ~ r1(X98, X103))) & (p3(X98) | ! [X104] : (~ p102(X104) | ~ p3(X104) | ~ r1(X98, X104))))) & (~ p103(X98) | ((~ p4(X98) | ! [X105] : (~ p103(X105) | p4(X105) | ~ r1(X98, X105))) & (p4(X98) | ! [X106] : (~ p103(X106) | ~ p4(X106) | ~ r1(X98, X106))))) & (~ p104(X98) | ((~ p5(X98) | ! [X107] : (~ p104(X107) | p5(X107) | ~ r1(X98, X107))) & (p5(X98) | ! [X108] : (~ p104(X108) | ~ p5(X108) | ~ r1(X98, X108))))) & (~ p105(X98) | ((~ p6(X98) | ! [X109] : (~ p105(X109) | p6(X109) | ~ r1(X98, X109))) & (p6(X98) | ! [X110] : (~ p105(X110) | ~ p6(X110) | ~ r1(X98, X110))))) & (~ p100(X98) | p101(X98) | (? [X111] : (p101(X111) & ~ p102(X111) & p2(X111) & r1(X98, X111)) & ? [X112] : (p101(X112) & ~ p102(X112) & ~ p2(X112) & r1(X98, X112)))) & (~ p101(X98) | p102(X98) | (? [X113] : (p102(X113) & ~ p103(X113) & p3(X113) & r1(X98, X113)) & ? [X114] : (p102(X114) & ~ p103(X114) & ~ p3(X114) & r1(X98, X114)))) & (~ p102(X98) | p103(X98) | (? [X115] : (p103(X115) & ~ p104(X115) & p4(X115) & r1(X98, X115)) & ? [X116] : (p103(X116) & ~ p104(X116) & ~ p4(X116) & r1(X98, X116)))) & (~ p103(X98) | p104(X98) | (? [X117] : (p104(X117) & ~ p105(X117) & p5(X117) & r1(X98, X117)) & ? [X118] : (p104(X118) & ~ p105(X118) & ~ p5(X118) & r1(X98, X118)))) & (~ p104(X98) | p105(X98) | (? [X119] : (p105(X119) & p6(X119) & r1(X98, X119)) & ? [X120] : (p105(X120) & ~ p6(X120) & r1(X98, X120))))) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X0, X95)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (((~ p101(X125) | p100(X125)) & (~ p102(X125) | p101(X125)) & (~ p103(X125) | p102(X125)) & (~ p104(X125) | p103(X125)) & (~ p105(X125) | p104(X125)) & (~ p100(X125) | ((~ p1(X125) | ! [X126] : (~ p100(X126) | p1(X126) | ~ r1(X125, X126))) & (p1(X125) | ! [X127] : (~ p100(X127) | ~ p1(X127) | ~ r1(X125, X127))))) & (~ p101(X125) | ((~ p2(X125) | ! [X128] : (~ p101(X128) | p2(X128) | ~ r1(X125, X128))) & (p2(X125) | ! [X129] : (~ p101(X129) | ~ p2(X129) | ~ r1(X125, X129))))) & (~ p102(X125) | ((~ p3(X125) | ! [X130] : (~ p102(X130) | p3(X130) | ~ r1(X125, X130))) & (p3(X125) | ! [X131] : (~ p102(X131) | ~ p3(X131) | ~ r1(X125, X131))))) & (~ p103(X125) | ((~ p4(X125) | ! [X132] : (~ p103(X132) | p4(X132) | ~ r1(X125, X132))) & (p4(X125) | ! [X133] : (~ p103(X133) | ~ p4(X133) | ~ r1(X125, X133))))) & (~ p104(X125) | ((~ p5(X125) | ! [X134] : (~ p104(X134) | p5(X134) | ~ r1(X125, X134))) & (p5(X125) | ! [X135] : (~ p104(X135) | ~ p5(X135) | ~ r1(X125, X135))))) & (~ p105(X125) | ((~ p6(X125) | ! [X136] : (~ p105(X136) | p6(X136) | ~ r1(X125, X136))) & (p6(X125) | ! [X137] : (~ p105(X137) | ~ p6(X137) | ~ r1(X125, X137))))) & (~ p100(X125) | p101(X125) | (? [X138] : (p101(X138) & ~ p102(X138) & p2(X138) & r1(X125, X138)) & ? [X139] : (p101(X139) & ~ p102(X139) & ~ p2(X139) & r1(X125, X139)))) & (~ p101(X125) | p102(X125) | (? [X140] : (p102(X140) & ~ p103(X140) & p3(X140) & r1(X125, X140)) & ? [X141] : (p102(X141) & ~ p103(X141) & ~ p3(X141) & r1(X125, X141)))) & (~ p102(X125) | p103(X125) | (? [X142] : (p103(X142) & ~ p104(X142) & p4(X142) & r1(X125, X142)) & ? [X143] : (p103(X143) & ~ p104(X143) & ~ p4(X143) & r1(X125, X143)))) & (~ p103(X125) | p104(X125) | (? [X144] : (p104(X144) & ~ p105(X144) & p5(X144) & r1(X125, X144)) & ? [X145] : (p104(X145) & ~ p105(X145) & ~ p5(X145) & r1(X125, X145)))) & (~ p104(X125) | p105(X125) | (? [X146] : (p105(X146) & p6(X146) & r1(X125, X146)) & ? [X147] : (p105(X147) & ~ p6(X147) & r1(X125, X147))))) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121)) & ! [X148] : (! [X149] : (! [X150] : (! [X151] : (! [X152] : (p3(X152) | ~ r1(X151, X152)) | ~ r1(X150, X151)) | ~ r1(X149, X150)) | ~ r1(X148, X149)) | ~ r1(X0, X148))), inference(flattening, [], [f6])).
fof(f6, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p102(X0) | ((~ p3(X0) | ! [X5] : (~ p102(X5) | p3(X5) | ~ r1(X0, X5))) & (p3(X0) | ! [X6] : (~ p102(X6) | ~ p3(X6) | ~ r1(X0, X6))))) & (~ p103(X0) | ((~ p4(X0) | ! [X7] : (~ p103(X7) | p4(X7) | ~ r1(X0, X7))) & (p4(X0) | ! [X8] : (~ p103(X8) | ~ p4(X8) | ~ r1(X0, X8))))) & (~ p104(X0) | ((~ p5(X0) | ! [X9] : (~ p104(X9) | p5(X9) | ~ r1(X0, X9))) & (p5(X0) | ! [X10] : (~ p104(X10) | ~ p5(X10) | ~ r1(X0, X10))))) & (~ p105(X0) | ((~ p6(X0) | ! [X11] : (~ p105(X11) | p6(X11) | ~ r1(X0, X11))) & (p6(X0) | ! [X12] : (~ p105(X12) | ~ p6(X12) | ~ r1(X0, X12))))) & ((~ p100(X0) | p101(X0)) | (? [X13] : ((p101(X13) & ~ p102(X13) & p2(X13)) & r1(X0, X13)) & ? [X14] : ((p101(X14) & ~ p102(X14) & ~ p2(X14)) & r1(X0, X14)))) & ((~ p101(X0) | p102(X0)) | (? [X15] : ((p102(X15) & ~ p103(X15) & p3(X15)) & r1(X0, X15)) & ? [X16] : ((p102(X16) & ~ p103(X16) & ~ p3(X16)) & r1(X0, X16)))) & ((~ p102(X0) | p103(X0)) | (? [X17] : ((p103(X17) & ~ p104(X17) & p4(X17)) & r1(X0, X17)) & ? [X18] : ((p103(X18) & ~ p104(X18) & ~ p4(X18)) & r1(X0, X18)))) & ((~ p103(X0) | p104(X0)) | (? [X19] : ((p104(X19) & ~ p105(X19) & p5(X19)) & r1(X0, X19)) & ? [X20] : ((p104(X20) & ~ p105(X20) & ~ p5(X20)) & r1(X0, X20)))) & ((~ p104(X0) | p105(X0)) | (? [X21] : ((p105(X21) & p6(X21)) & r1(X0, X21)) & ? [X22] : ((p105(X22) & ~ p6(X22)) & r1(X0, X22)))) & ! [X23] : (((~ p101(X23) | p100(X23)) & (~ p102(X23) | p101(X23)) & (~ p103(X23) | p102(X23)) & (~ p104(X23) | p103(X23)) & (~ p105(X23) | p104(X23)) & (~ p100(X23) | ((~ p1(X23) | ! [X24] : (~ p100(X24) | p1(X24) | ~ r1(X23, X24))) & (p1(X23) | ! [X25] : (~ p100(X25) | ~ p1(X25) | ~ r1(X23, X25))))) & (~ p101(X23) | ((~ p2(X23) | ! [X26] : (~ p101(X26) | p2(X26) | ~ r1(X23, X26))) & (p2(X23) | ! [X27] : (~ p101(X27) | ~ p2(X27) | ~ r1(X23, X27))))) & (~ p102(X23) | ((~ p3(X23) | ! [X28] : (~ p102(X28) | p3(X28) | ~ r1(X23, X28))) & (p3(X23) | ! [X29] : (~ p102(X29) | ~ p3(X29) | ~ r1(X23, X29))))) & (~ p103(X23) | ((~ p4(X23) | ! [X30] : (~ p103(X30) | p4(X30) | ~ r1(X23, X30))) & (p4(X23) | ! [X31] : (~ p103(X31) | ~ p4(X31) | ~ r1(X23, X31))))) & (~ p104(X23) | ((~ p5(X23) | ! [X32] : (~ p104(X32) | p5(X32) | ~ r1(X23, X32))) & (p5(X23) | ! [X33] : (~ p104(X33) | ~ p5(X33) | ~ r1(X23, X33))))) & (~ p105(X23) | ((~ p6(X23) | ! [X34] : (~ p105(X34) | p6(X34) | ~ r1(X23, X34))) & (p6(X23) | ! [X35] : (~ p105(X35) | ~ p6(X35) | ~ r1(X23, X35))))) & ((~ p100(X23) | p101(X23)) | (? [X36] : ((p101(X36) & ~ p102(X36) & p2(X36)) & r1(X23, X36)) & ? [X37] : ((p101(X37) & ~ p102(X37) & ~ p2(X37)) & r1(X23, X37)))) & ((~ p101(X23) | p102(X23)) | (? [X38] : ((p102(X38) & ~ p103(X38) & p3(X38)) & r1(X23, X38)) & ? [X39] : ((p102(X39) & ~ p103(X39) & ~ p3(X39)) & r1(X23, X39)))) & ((~ p102(X23) | p103(X23)) | (? [X40] : ((p103(X40) & ~ p104(X40) & p4(X40)) & r1(X23, X40)) & ? [X41] : ((p103(X41) & ~ p104(X41) & ~ p4(X41)) & r1(X23, X41)))) & ((~ p103(X23) | p104(X23)) | (? [X42] : ((p104(X42) & ~ p105(X42) & p5(X42)) & r1(X23, X42)) & ? [X43] : ((p104(X43) & ~ p105(X43) & ~ p5(X43)) & r1(X23, X43)))) & ((~ p104(X23) | p105(X23)) | (? [X44] : ((p105(X44) & p6(X44)) & r1(X23, X44)) & ? [X45] : ((p105(X45) & ~ p6(X45)) & r1(X23, X45))))) | ~ r1(X0, X23)) & ! [X46] : (! [X47] : (((~ p101(X47) | p100(X47)) & (~ p102(X47) | p101(X47)) & (~ p103(X47) | p102(X47)) & (~ p104(X47) | p103(X47)) & (~ p105(X47) | p104(X47)) & (~ p100(X47) | ((~ p1(X47) | ! [X48] : (~ p100(X48) | p1(X48) | ~ r1(X47, X48))) & (p1(X47) | ! [X49] : (~ p100(X49) | ~ p1(X49) | ~ r1(X47, X49))))) & (~ p101(X47) | ((~ p2(X47) | ! [X50] : (~ p101(X50) | p2(X50) | ~ r1(X47, X50))) & (p2(X47) | ! [X51] : (~ p101(X51) | ~ p2(X51) | ~ r1(X47, X51))))) & (~ p102(X47) | ((~ p3(X47) | ! [X52] : (~ p102(X52) | p3(X52) | ~ r1(X47, X52))) & (p3(X47) | ! [X53] : (~ p102(X53) | ~ p3(X53) | ~ r1(X47, X53))))) & (~ p103(X47) | ((~ p4(X47) | ! [X54] : (~ p103(X54) | p4(X54) | ~ r1(X47, X54))) & (p4(X47) | ! [X55] : (~ p103(X55) | ~ p4(X55) | ~ r1(X47, X55))))) & (~ p104(X47) | ((~ p5(X47) | ! [X56] : (~ p104(X56) | p5(X56) | ~ r1(X47, X56))) & (p5(X47) | ! [X57] : (~ p104(X57) | ~ p5(X57) | ~ r1(X47, X57))))) & (~ p105(X47) | ((~ p6(X47) | ! [X58] : (~ p105(X58) | p6(X58) | ~ r1(X47, X58))) & (p6(X47) | ! [X59] : (~ p105(X59) | ~ p6(X59) | ~ r1(X47, X59))))) & ((~ p100(X47) | p101(X47)) | (? [X60] : ((p101(X60) & ~ p102(X60) & p2(X60)) & r1(X47, X60)) & ? [X61] : ((p101(X61) & ~ p102(X61) & ~ p2(X61)) & r1(X47, X61)))) & ((~ p101(X47) | p102(X47)) | (? [X62] : ((p102(X62) & ~ p103(X62) & p3(X62)) & r1(X47, X62)) & ? [X63] : ((p102(X63) & ~ p103(X63) & ~ p3(X63)) & r1(X47, X63)))) & ((~ p102(X47) | p103(X47)) | (? [X64] : ((p103(X64) & ~ p104(X64) & p4(X64)) & r1(X47, X64)) & ? [X65] : ((p103(X65) & ~ p104(X65) & ~ p4(X65)) & r1(X47, X65)))) & ((~ p103(X47) | p104(X47)) | (? [X66] : ((p104(X66) & ~ p105(X66) & p5(X66)) & r1(X47, X66)) & ? [X67] : ((p104(X67) & ~ p105(X67) & ~ p5(X67)) & r1(X47, X67)))) & ((~ p104(X47) | p105(X47)) | (? [X68] : ((p105(X68) & p6(X68)) & r1(X47, X68)) & ? [X69] : ((p105(X69) & ~ p6(X69)) & r1(X47, X69))))) | ~ r1(X46, X47)) | ~ r1(X0, X46)) & ! [X70] : (! [X71] : (! [X72] : (((~ p101(X72) | p100(X72)) & (~ p102(X72) | p101(X72)) & (~ p103(X72) | p102(X72)) & (~ p104(X72) | p103(X72)) & (~ p105(X72) | p104(X72)) & (~ p100(X72) | ((~ p1(X72) | ! [X73] : (~ p100(X73) | p1(X73) | ~ r1(X72, X73))) & (p1(X72) | ! [X74] : (~ p100(X74) | ~ p1(X74) | ~ r1(X72, X74))))) & (~ p101(X72) | ((~ p2(X72) | ! [X75] : (~ p101(X75) | p2(X75) | ~ r1(X72, X75))) & (p2(X72) | ! [X76] : (~ p101(X76) | ~ p2(X76) | ~ r1(X72, X76))))) & (~ p102(X72) | ((~ p3(X72) | ! [X77] : (~ p102(X77) | p3(X77) | ~ r1(X72, X77))) & (p3(X72) | ! [X78] : (~ p102(X78) | ~ p3(X78) | ~ r1(X72, X78))))) & (~ p103(X72) | ((~ p4(X72) | ! [X79] : (~ p103(X79) | p4(X79) | ~ r1(X72, X79))) & (p4(X72) | ! [X80] : (~ p103(X80) | ~ p4(X80) | ~ r1(X72, X80))))) & (~ p104(X72) | ((~ p5(X72) | ! [X81] : (~ p104(X81) | p5(X81) | ~ r1(X72, X81))) & (p5(X72) | ! [X82] : (~ p104(X82) | ~ p5(X82) | ~ r1(X72, X82))))) & (~ p105(X72) | ((~ p6(X72) | ! [X83] : (~ p105(X83) | p6(X83) | ~ r1(X72, X83))) & (p6(X72) | ! [X84] : (~ p105(X84) | ~ p6(X84) | ~ r1(X72, X84))))) & ((~ p100(X72) | p101(X72)) | (? [X85] : ((p101(X85) & ~ p102(X85) & p2(X85)) & r1(X72, X85)) & ? [X86] : ((p101(X86) & ~ p102(X86) & ~ p2(X86)) & r1(X72, X86)))) & ((~ p101(X72) | p102(X72)) | (? [X87] : ((p102(X87) & ~ p103(X87) & p3(X87)) & r1(X72, X87)) & ? [X88] : ((p102(X88) & ~ p103(X88) & ~ p3(X88)) & r1(X72, X88)))) & ((~ p102(X72) | p103(X72)) | (? [X89] : ((p103(X89) & ~ p104(X89) & p4(X89)) & r1(X72, X89)) & ? [X90] : ((p103(X90) & ~ p104(X90) & ~ p4(X90)) & r1(X72, X90)))) & ((~ p103(X72) | p104(X72)) | (? [X91] : ((p104(X91) & ~ p105(X91) & p5(X91)) & r1(X72, X91)) & ? [X92] : ((p104(X92) & ~ p105(X92) & ~ p5(X92)) & r1(X72, X92)))) & ((~ p104(X72) | p105(X72)) | (? [X93] : ((p105(X93) & p6(X93)) & r1(X72, X93)) & ? [X94] : ((p105(X94) & ~ p6(X94)) & r1(X72, X94))))) | ~ r1(X71, X72)) | ~ r1(X70, X71)) | ~ r1(X0, X70)) & ! [X95] : (! [X96] : (! [X97] : (! [X98] : (((~ p101(X98) | p100(X98)) & (~ p102(X98) | p101(X98)) & (~ p103(X98) | p102(X98)) & (~ p104(X98) | p103(X98)) & (~ p105(X98) | p104(X98)) & (~ p100(X98) | ((~ p1(X98) | ! [X99] : (~ p100(X99) | p1(X99) | ~ r1(X98, X99))) & (p1(X98) | ! [X100] : (~ p100(X100) | ~ p1(X100) | ~ r1(X98, X100))))) & (~ p101(X98) | ((~ p2(X98) | ! [X101] : (~ p101(X101) | p2(X101) | ~ r1(X98, X101))) & (p2(X98) | ! [X102] : (~ p101(X102) | ~ p2(X102) | ~ r1(X98, X102))))) & (~ p102(X98) | ((~ p3(X98) | ! [X103] : (~ p102(X103) | p3(X103) | ~ r1(X98, X103))) & (p3(X98) | ! [X104] : (~ p102(X104) | ~ p3(X104) | ~ r1(X98, X104))))) & (~ p103(X98) | ((~ p4(X98) | ! [X105] : (~ p103(X105) | p4(X105) | ~ r1(X98, X105))) & (p4(X98) | ! [X106] : (~ p103(X106) | ~ p4(X106) | ~ r1(X98, X106))))) & (~ p104(X98) | ((~ p5(X98) | ! [X107] : (~ p104(X107) | p5(X107) | ~ r1(X98, X107))) & (p5(X98) | ! [X108] : (~ p104(X108) | ~ p5(X108) | ~ r1(X98, X108))))) & (~ p105(X98) | ((~ p6(X98) | ! [X109] : (~ p105(X109) | p6(X109) | ~ r1(X98, X109))) & (p6(X98) | ! [X110] : (~ p105(X110) | ~ p6(X110) | ~ r1(X98, X110))))) & ((~ p100(X98) | p101(X98)) | (? [X111] : ((p101(X111) & ~ p102(X111) & p2(X111)) & r1(X98, X111)) & ? [X112] : ((p101(X112) & ~ p102(X112) & ~ p2(X112)) & r1(X98, X112)))) & ((~ p101(X98) | p102(X98)) | (? [X113] : ((p102(X113) & ~ p103(X113) & p3(X113)) & r1(X98, X113)) & ? [X114] : ((p102(X114) & ~ p103(X114) & ~ p3(X114)) & r1(X98, X114)))) & ((~ p102(X98) | p103(X98)) | (? [X115] : ((p103(X115) & ~ p104(X115) & p4(X115)) & r1(X98, X115)) & ? [X116] : ((p103(X116) & ~ p104(X116) & ~ p4(X116)) & r1(X98, X116)))) & ((~ p103(X98) | p104(X98)) | (? [X117] : ((p104(X117) & ~ p105(X117) & p5(X117)) & r1(X98, X117)) & ? [X118] : ((p104(X118) & ~ p105(X118) & ~ p5(X118)) & r1(X98, X118)))) & ((~ p104(X98) | p105(X98)) | (? [X119] : ((p105(X119) & p6(X119)) & r1(X98, X119)) & ? [X120] : ((p105(X120) & ~ p6(X120)) & r1(X98, X120))))) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X0, X95)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (((~ p101(X125) | p100(X125)) & (~ p102(X125) | p101(X125)) & (~ p103(X125) | p102(X125)) & (~ p104(X125) | p103(X125)) & (~ p105(X125) | p104(X125)) & (~ p100(X125) | ((~ p1(X125) | ! [X126] : (~ p100(X126) | p1(X126) | ~ r1(X125, X126))) & (p1(X125) | ! [X127] : (~ p100(X127) | ~ p1(X127) | ~ r1(X125, X127))))) & (~ p101(X125) | ((~ p2(X125) | ! [X128] : (~ p101(X128) | p2(X128) | ~ r1(X125, X128))) & (p2(X125) | ! [X129] : (~ p101(X129) | ~ p2(X129) | ~ r1(X125, X129))))) & (~ p102(X125) | ((~ p3(X125) | ! [X130] : (~ p102(X130) | p3(X130) | ~ r1(X125, X130))) & (p3(X125) | ! [X131] : (~ p102(X131) | ~ p3(X131) | ~ r1(X125, X131))))) & (~ p103(X125) | ((~ p4(X125) | ! [X132] : (~ p103(X132) | p4(X132) | ~ r1(X125, X132))) & (p4(X125) | ! [X133] : (~ p103(X133) | ~ p4(X133) | ~ r1(X125, X133))))) & (~ p104(X125) | ((~ p5(X125) | ! [X134] : (~ p104(X134) | p5(X134) | ~ r1(X125, X134))) & (p5(X125) | ! [X135] : (~ p104(X135) | ~ p5(X135) | ~ r1(X125, X135))))) & (~ p105(X125) | ((~ p6(X125) | ! [X136] : (~ p105(X136) | p6(X136) | ~ r1(X125, X136))) & (p6(X125) | ! [X137] : (~ p105(X137) | ~ p6(X137) | ~ r1(X125, X137))))) & ((~ p100(X125) | p101(X125)) | (? [X138] : ((p101(X138) & ~ p102(X138) & p2(X138)) & r1(X125, X138)) & ? [X139] : ((p101(X139) & ~ p102(X139) & ~ p2(X139)) & r1(X125, X139)))) & ((~ p101(X125) | p102(X125)) | (? [X140] : ((p102(X140) & ~ p103(X140) & p3(X140)) & r1(X125, X140)) & ? [X141] : ((p102(X141) & ~ p103(X141) & ~ p3(X141)) & r1(X125, X141)))) & ((~ p102(X125) | p103(X125)) | (? [X142] : ((p103(X142) & ~ p104(X142) & p4(X142)) & r1(X125, X142)) & ? [X143] : ((p103(X143) & ~ p104(X143) & ~ p4(X143)) & r1(X125, X143)))) & ((~ p103(X125) | p104(X125)) | (? [X144] : ((p104(X144) & ~ p105(X144) & p5(X144)) & r1(X125, X144)) & ? [X145] : ((p104(X145) & ~ p105(X145) & ~ p5(X145)) & r1(X125, X145)))) & ((~ p104(X125) | p105(X125)) | (? [X146] : ((p105(X146) & p6(X146)) & r1(X125, X146)) & ? [X147] : ((p105(X147) & ~ p6(X147)) & r1(X125, X147))))) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121))) & ! [X148] : (! [X149] : (! [X150] : (! [X151] : (! [X152] : (p3(X152) | ~ r1(X151, X152)) | ~ r1(X150, X151)) | ~ r1(X149, X150)) | ~ r1(X148, X149)) | ~ r1(X0, X148))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p102(X0) | ((~ p3(X0) | ! [X5] : (~ p102(X5) | p3(X5) | ~ r1(X0, X5))) & (p3(X0) | ! [X6] : (~ p102(X6) | ~ p3(X6) | ~ r1(X0, X6))))) & (~ p103(X0) | ((~ p4(X0) | ! [X7] : (~ p103(X7) | p4(X7) | ~ r1(X0, X7))) & (p4(X0) | ! [X8] : (~ p103(X8) | ~ p4(X8) | ~ r1(X0, X8))))) & (~ p104(X0) | ((~ p5(X0) | ! [X9] : (~ p104(X9) | p5(X9) | ~ r1(X0, X9))) & (p5(X0) | ! [X10] : (~ p104(X10) | ~ p5(X10) | ~ r1(X0, X10))))) & (~ p105(X0) | ((~ p6(X0) | ! [X11] : (~ p105(X11) | p6(X11) | ~ r1(X0, X11))) & (p6(X0) | ! [X12] : (~ p105(X12) | ~ p6(X12) | ~ r1(X0, X12))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X13] : (~ (p101(X13) & ~ p102(X13) & p2(X13)) | ~ r1(X0, X13)) & ~ ! [X14] : (~ (p101(X14) & ~ p102(X14) & ~ p2(X14)) | ~ r1(X0, X14)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X15] : (~ (p102(X15) & ~ p103(X15) & p3(X15)) | ~ r1(X0, X15)) & ~ ! [X16] : (~ (p102(X16) & ~ p103(X16) & ~ p3(X16)) | ~ r1(X0, X16)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X17] : (~ (p103(X17) & ~ p104(X17) & p4(X17)) | ~ r1(X0, X17)) & ~ ! [X18] : (~ (p103(X18) & ~ p104(X18) & ~ p4(X18)) | ~ r1(X0, X18)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X19] : (~ (p104(X19) & ~ p105(X19) & p5(X19)) | ~ r1(X0, X19)) & ~ ! [X20] : (~ (p104(X20) & ~ p105(X20) & ~ p5(X20)) | ~ r1(X0, X20)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X21] : (~ (p105(X21) & p6(X21)) | ~ r1(X0, X21)) & ~ ! [X22] : (~ (p105(X22) & ~ p6(X22)) | ~ r1(X0, X22)))) & ! [X23] : (((~ p101(X23) | p100(X23)) & (~ p102(X23) | p101(X23)) & (~ p103(X23) | p102(X23)) & (~ p104(X23) | p103(X23)) & (~ p105(X23) | p104(X23)) & (~ p100(X23) | ((~ p1(X23) | ! [X24] : (~ p100(X24) | p1(X24) | ~ r1(X23, X24))) & (p1(X23) | ! [X25] : (~ p100(X25) | ~ p1(X25) | ~ r1(X23, X25))))) & (~ p101(X23) | ((~ p2(X23) | ! [X26] : (~ p101(X26) | p2(X26) | ~ r1(X23, X26))) & (p2(X23) | ! [X27] : (~ p101(X27) | ~ p2(X27) | ~ r1(X23, X27))))) & (~ p102(X23) | ((~ p3(X23) | ! [X28] : (~ p102(X28) | p3(X28) | ~ r1(X23, X28))) & (p3(X23) | ! [X29] : (~ p102(X29) | ~ p3(X29) | ~ r1(X23, X29))))) & (~ p103(X23) | ((~ p4(X23) | ! [X30] : (~ p103(X30) | p4(X30) | ~ r1(X23, X30))) & (p4(X23) | ! [X31] : (~ p103(X31) | ~ p4(X31) | ~ r1(X23, X31))))) & (~ p104(X23) | ((~ p5(X23) | ! [X32] : (~ p104(X32) | p5(X32) | ~ r1(X23, X32))) & (p5(X23) | ! [X33] : (~ p104(X33) | ~ p5(X33) | ~ r1(X23, X33))))) & (~ p105(X23) | ((~ p6(X23) | ! [X34] : (~ p105(X34) | p6(X34) | ~ r1(X23, X34))) & (p6(X23) | ! [X35] : (~ p105(X35) | ~ p6(X35) | ~ r1(X23, X35))))) & (~ (p100(X23) & ~ p101(X23)) | (~ ! [X36] : (~ (p101(X36) & ~ p102(X36) & p2(X36)) | ~ r1(X23, X36)) & ~ ! [X37] : (~ (p101(X37) & ~ p102(X37) & ~ p2(X37)) | ~ r1(X23, X37)))) & (~ (p101(X23) & ~ p102(X23)) | (~ ! [X38] : (~ (p102(X38) & ~ p103(X38) & p3(X38)) | ~ r1(X23, X38)) & ~ ! [X39] : (~ (p102(X39) & ~ p103(X39) & ~ p3(X39)) | ~ r1(X23, X39)))) & (~ (p102(X23) & ~ p103(X23)) | (~ ! [X40] : (~ (p103(X40) & ~ p104(X40) & p4(X40)) | ~ r1(X23, X40)) & ~ ! [X41] : (~ (p103(X41) & ~ p104(X41) & ~ p4(X41)) | ~ r1(X23, X41)))) & (~ (p103(X23) & ~ p104(X23)) | (~ ! [X42] : (~ (p104(X42) & ~ p105(X42) & p5(X42)) | ~ r1(X23, X42)) & ~ ! [X43] : (~ (p104(X43) & ~ p105(X43) & ~ p5(X43)) | ~ r1(X23, X43)))) & (~ (p104(X23) & ~ p105(X23)) | (~ ! [X44] : (~ (p105(X44) & p6(X44)) | ~ r1(X23, X44)) & ~ ! [X45] : (~ (p105(X45) & ~ p6(X45)) | ~ r1(X23, X45))))) | ~ r1(X0, X23)) & ! [X46] : (! [X47] : (((~ p101(X47) | p100(X47)) & (~ p102(X47) | p101(X47)) & (~ p103(X47) | p102(X47)) & (~ p104(X47) | p103(X47)) & (~ p105(X47) | p104(X47)) & (~ p100(X47) | ((~ p1(X47) | ! [X48] : (~ p100(X48) | p1(X48) | ~ r1(X47, X48))) & (p1(X47) | ! [X49] : (~ p100(X49) | ~ p1(X49) | ~ r1(X47, X49))))) & (~ p101(X47) | ((~ p2(X47) | ! [X50] : (~ p101(X50) | p2(X50) | ~ r1(X47, X50))) & (p2(X47) | ! [X51] : (~ p101(X51) | ~ p2(X51) | ~ r1(X47, X51))))) & (~ p102(X47) | ((~ p3(X47) | ! [X52] : (~ p102(X52) | p3(X52) | ~ r1(X47, X52))) & (p3(X47) | ! [X53] : (~ p102(X53) | ~ p3(X53) | ~ r1(X47, X53))))) & (~ p103(X47) | ((~ p4(X47) | ! [X54] : (~ p103(X54) | p4(X54) | ~ r1(X47, X54))) & (p4(X47) | ! [X55] : (~ p103(X55) | ~ p4(X55) | ~ r1(X47, X55))))) & (~ p104(X47) | ((~ p5(X47) | ! [X56] : (~ p104(X56) | p5(X56) | ~ r1(X47, X56))) & (p5(X47) | ! [X57] : (~ p104(X57) | ~ p5(X57) | ~ r1(X47, X57))))) & (~ p105(X47) | ((~ p6(X47) | ! [X58] : (~ p105(X58) | p6(X58) | ~ r1(X47, X58))) & (p6(X47) | ! [X59] : (~ p105(X59) | ~ p6(X59) | ~ r1(X47, X59))))) & (~ (p100(X47) & ~ p101(X47)) | (~ ! [X60] : (~ (p101(X60) & ~ p102(X60) & p2(X60)) | ~ r1(X47, X60)) & ~ ! [X61] : (~ (p101(X61) & ~ p102(X61) & ~ p2(X61)) | ~ r1(X47, X61)))) & (~ (p101(X47) & ~ p102(X47)) | (~ ! [X62] : (~ (p102(X62) & ~ p103(X62) & p3(X62)) | ~ r1(X47, X62)) & ~ ! [X63] : (~ (p102(X63) & ~ p103(X63) & ~ p3(X63)) | ~ r1(X47, X63)))) & (~ (p102(X47) & ~ p103(X47)) | (~ ! [X64] : (~ (p103(X64) & ~ p104(X64) & p4(X64)) | ~ r1(X47, X64)) & ~ ! [X65] : (~ (p103(X65) & ~ p104(X65) & ~ p4(X65)) | ~ r1(X47, X65)))) & (~ (p103(X47) & ~ p104(X47)) | (~ ! [X66] : (~ (p104(X66) & ~ p105(X66) & p5(X66)) | ~ r1(X47, X66)) & ~ ! [X67] : (~ (p104(X67) & ~ p105(X67) & ~ p5(X67)) | ~ r1(X47, X67)))) & (~ (p104(X47) & ~ p105(X47)) | (~ ! [X68] : (~ (p105(X68) & p6(X68)) | ~ r1(X47, X68)) & ~ ! [X69] : (~ (p105(X69) & ~ p6(X69)) | ~ r1(X47, X69))))) | ~ r1(X46, X47)) | ~ r1(X0, X46)) & ! [X70] : (! [X71] : (! [X72] : (((~ p101(X72) | p100(X72)) & (~ p102(X72) | p101(X72)) & (~ p103(X72) | p102(X72)) & (~ p104(X72) | p103(X72)) & (~ p105(X72) | p104(X72)) & (~ p100(X72) | ((~ p1(X72) | ! [X73] : (~ p100(X73) | p1(X73) | ~ r1(X72, X73))) & (p1(X72) | ! [X74] : (~ p100(X74) | ~ p1(X74) | ~ r1(X72, X74))))) & (~ p101(X72) | ((~ p2(X72) | ! [X75] : (~ p101(X75) | p2(X75) | ~ r1(X72, X75))) & (p2(X72) | ! [X76] : (~ p101(X76) | ~ p2(X76) | ~ r1(X72, X76))))) & (~ p102(X72) | ((~ p3(X72) | ! [X77] : (~ p102(X77) | p3(X77) | ~ r1(X72, X77))) & (p3(X72) | ! [X78] : (~ p102(X78) | ~ p3(X78) | ~ r1(X72, X78))))) & (~ p103(X72) | ((~ p4(X72) | ! [X79] : (~ p103(X79) | p4(X79) | ~ r1(X72, X79))) & (p4(X72) | ! [X80] : (~ p103(X80) | ~ p4(X80) | ~ r1(X72, X80))))) & (~ p104(X72) | ((~ p5(X72) | ! [X81] : (~ p104(X81) | p5(X81) | ~ r1(X72, X81))) & (p5(X72) | ! [X82] : (~ p104(X82) | ~ p5(X82) | ~ r1(X72, X82))))) & (~ p105(X72) | ((~ p6(X72) | ! [X83] : (~ p105(X83) | p6(X83) | ~ r1(X72, X83))) & (p6(X72) | ! [X84] : (~ p105(X84) | ~ p6(X84) | ~ r1(X72, X84))))) & (~ (p100(X72) & ~ p101(X72)) | (~ ! [X85] : (~ (p101(X85) & ~ p102(X85) & p2(X85)) | ~ r1(X72, X85)) & ~ ! [X86] : (~ (p101(X86) & ~ p102(X86) & ~ p2(X86)) | ~ r1(X72, X86)))) & (~ (p101(X72) & ~ p102(X72)) | (~ ! [X87] : (~ (p102(X87) & ~ p103(X87) & p3(X87)) | ~ r1(X72, X87)) & ~ ! [X88] : (~ (p102(X88) & ~ p103(X88) & ~ p3(X88)) | ~ r1(X72, X88)))) & (~ (p102(X72) & ~ p103(X72)) | (~ ! [X89] : (~ (p103(X89) & ~ p104(X89) & p4(X89)) | ~ r1(X72, X89)) & ~ ! [X90] : (~ (p103(X90) & ~ p104(X90) & ~ p4(X90)) | ~ r1(X72, X90)))) & (~ (p103(X72) & ~ p104(X72)) | (~ ! [X91] : (~ (p104(X91) & ~ p105(X91) & p5(X91)) | ~ r1(X72, X91)) & ~ ! [X92] : (~ (p104(X92) & ~ p105(X92) & ~ p5(X92)) | ~ r1(X72, X92)))) & (~ (p104(X72) & ~ p105(X72)) | (~ ! [X93] : (~ (p105(X93) & p6(X93)) | ~ r1(X72, X93)) & ~ ! [X94] : (~ (p105(X94) & ~ p6(X94)) | ~ r1(X72, X94))))) | ~ r1(X71, X72)) | ~ r1(X70, X71)) | ~ r1(X0, X70)) & ! [X95] : (! [X96] : (! [X97] : (! [X98] : (((~ p101(X98) | p100(X98)) & (~ p102(X98) | p101(X98)) & (~ p103(X98) | p102(X98)) & (~ p104(X98) | p103(X98)) & (~ p105(X98) | p104(X98)) & (~ p100(X98) | ((~ p1(X98) | ! [X99] : (~ p100(X99) | p1(X99) | ~ r1(X98, X99))) & (p1(X98) | ! [X100] : (~ p100(X100) | ~ p1(X100) | ~ r1(X98, X100))))) & (~ p101(X98) | ((~ p2(X98) | ! [X101] : (~ p101(X101) | p2(X101) | ~ r1(X98, X101))) & (p2(X98) | ! [X102] : (~ p101(X102) | ~ p2(X102) | ~ r1(X98, X102))))) & (~ p102(X98) | ((~ p3(X98) | ! [X103] : (~ p102(X103) | p3(X103) | ~ r1(X98, X103))) & (p3(X98) | ! [X104] : (~ p102(X104) | ~ p3(X104) | ~ r1(X98, X104))))) & (~ p103(X98) | ((~ p4(X98) | ! [X105] : (~ p103(X105) | p4(X105) | ~ r1(X98, X105))) & (p4(X98) | ! [X106] : (~ p103(X106) | ~ p4(X106) | ~ r1(X98, X106))))) & (~ p104(X98) | ((~ p5(X98) | ! [X107] : (~ p104(X107) | p5(X107) | ~ r1(X98, X107))) & (p5(X98) | ! [X108] : (~ p104(X108) | ~ p5(X108) | ~ r1(X98, X108))))) & (~ p105(X98) | ((~ p6(X98) | ! [X109] : (~ p105(X109) | p6(X109) | ~ r1(X98, X109))) & (p6(X98) | ! [X110] : (~ p105(X110) | ~ p6(X110) | ~ r1(X98, X110))))) & (~ (p100(X98) & ~ p101(X98)) | (~ ! [X111] : (~ (p101(X111) & ~ p102(X111) & p2(X111)) | ~ r1(X98, X111)) & ~ ! [X112] : (~ (p101(X112) & ~ p102(X112) & ~ p2(X112)) | ~ r1(X98, X112)))) & (~ (p101(X98) & ~ p102(X98)) | (~ ! [X113] : (~ (p102(X113) & ~ p103(X113) & p3(X113)) | ~ r1(X98, X113)) & ~ ! [X114] : (~ (p102(X114) & ~ p103(X114) & ~ p3(X114)) | ~ r1(X98, X114)))) & (~ (p102(X98) & ~ p103(X98)) | (~ ! [X115] : (~ (p103(X115) & ~ p104(X115) & p4(X115)) | ~ r1(X98, X115)) & ~ ! [X116] : (~ (p103(X116) & ~ p104(X116) & ~ p4(X116)) | ~ r1(X98, X116)))) & (~ (p103(X98) & ~ p104(X98)) | (~ ! [X117] : (~ (p104(X117) & ~ p105(X117) & p5(X117)) | ~ r1(X98, X117)) & ~ ! [X118] : (~ (p104(X118) & ~ p105(X118) & ~ p5(X118)) | ~ r1(X98, X118)))) & (~ (p104(X98) & ~ p105(X98)) | (~ ! [X119] : (~ (p105(X119) & p6(X119)) | ~ r1(X98, X119)) & ~ ! [X120] : (~ (p105(X120) & ~ p6(X120)) | ~ r1(X98, X120))))) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X0, X95)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (((~ p101(X125) | p100(X125)) & (~ p102(X125) | p101(X125)) & (~ p103(X125) | p102(X125)) & (~ p104(X125) | p103(X125)) & (~ p105(X125) | p104(X125)) & (~ p100(X125) | ((~ p1(X125) | ! [X126] : (~ p100(X126) | p1(X126) | ~ r1(X125, X126))) & (p1(X125) | ! [X127] : (~ p100(X127) | ~ p1(X127) | ~ r1(X125, X127))))) & (~ p101(X125) | ((~ p2(X125) | ! [X128] : (~ p101(X128) | p2(X128) | ~ r1(X125, X128))) & (p2(X125) | ! [X129] : (~ p101(X129) | ~ p2(X129) | ~ r1(X125, X129))))) & (~ p102(X125) | ((~ p3(X125) | ! [X130] : (~ p102(X130) | p3(X130) | ~ r1(X125, X130))) & (p3(X125) | ! [X131] : (~ p102(X131) | ~ p3(X131) | ~ r1(X125, X131))))) & (~ p103(X125) | ((~ p4(X125) | ! [X132] : (~ p103(X132) | p4(X132) | ~ r1(X125, X132))) & (p4(X125) | ! [X133] : (~ p103(X133) | ~ p4(X133) | ~ r1(X125, X133))))) & (~ p104(X125) | ((~ p5(X125) | ! [X134] : (~ p104(X134) | p5(X134) | ~ r1(X125, X134))) & (p5(X125) | ! [X135] : (~ p104(X135) | ~ p5(X135) | ~ r1(X125, X135))))) & (~ p105(X125) | ((~ p6(X125) | ! [X136] : (~ p105(X136) | p6(X136) | ~ r1(X125, X136))) & (p6(X125) | ! [X137] : (~ p105(X137) | ~ p6(X137) | ~ r1(X125, X137))))) & (~ (p100(X125) & ~ p101(X125)) | (~ ! [X138] : (~ (p101(X138) & ~ p102(X138) & p2(X138)) | ~ r1(X125, X138)) & ~ ! [X139] : (~ (p101(X139) & ~ p102(X139) & ~ p2(X139)) | ~ r1(X125, X139)))) & (~ (p101(X125) & ~ p102(X125)) | (~ ! [X140] : (~ (p102(X140) & ~ p103(X140) & p3(X140)) | ~ r1(X125, X140)) & ~ ! [X141] : (~ (p102(X141) & ~ p103(X141) & ~ p3(X141)) | ~ r1(X125, X141)))) & (~ (p102(X125) & ~ p103(X125)) | (~ ! [X142] : (~ (p103(X142) & ~ p104(X142) & p4(X142)) | ~ r1(X125, X142)) & ~ ! [X143] : (~ (p103(X143) & ~ p104(X143) & ~ p4(X143)) | ~ r1(X125, X143)))) & (~ (p103(X125) & ~ p104(X125)) | (~ ! [X144] : (~ (p104(X144) & ~ p105(X144) & p5(X144)) | ~ r1(X125, X144)) & ~ ! [X145] : (~ (p104(X145) & ~ p105(X145) & ~ p5(X145)) | ~ r1(X125, X145)))) & (~ (p104(X125) & ~ p105(X125)) | (~ ! [X146] : (~ (p105(X146) & p6(X146)) | ~ r1(X125, X146)) & ~ ! [X147] : (~ (p105(X147) & ~ p6(X147)) | ~ r1(X125, X147))))) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121))) | ~ ! [X148] : (! [X149] : (! [X150] : (! [X151] : (! [X152] : (p3(X152) | ~ r1(X151, X152)) | ~ r1(X150, X151)) | ~ r1(X149, X150)) | ~ r1(X148, X149)) | ~ r1(X0, X148))), inference(pure_predicate_removal, [], [f4])).
fof(f4, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p102(X0) | ((~ p3(X0) | ! [X5] : (~ p102(X5) | p3(X5) | ~ r1(X0, X5))) & (p3(X0) | ! [X6] : (~ p102(X6) | ~ p3(X6) | ~ r1(X0, X6))))) & (~ p103(X0) | ((~ p4(X0) | ! [X7] : (~ p103(X7) | p4(X7) | ~ r1(X0, X7))) & (p4(X0) | ! [X8] : (~ p103(X8) | ~ p4(X8) | ~ r1(X0, X8))))) & (~ p104(X0) | ((~ p5(X0) | ! [X9] : (~ p104(X9) | p5(X9) | ~ r1(X0, X9))) & (p5(X0) | ! [X10] : (~ p104(X10) | ~ p5(X10) | ~ r1(X0, X10))))) & (~ p105(X0) | ((~ p6(X0) | ! [X11] : (~ p105(X11) | p6(X11) | ~ r1(X0, X11))) & (p6(X0) | ! [X12] : (~ p105(X12) | ~ p6(X12) | ~ r1(X0, X12))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X13] : (~ (p101(X13) & ~ p102(X13) & p2(X13)) | ~ r1(X0, X13)) & ~ ! [X14] : (~ (p101(X14) & ~ p102(X14) & ~ p2(X14)) | ~ r1(X0, X14)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X15] : (~ (p102(X15) & ~ p103(X15) & p3(X15)) | ~ r1(X0, X15)) & ~ ! [X16] : (~ (p102(X16) & ~ p103(X16) & ~ p3(X16)) | ~ r1(X0, X16)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X17] : (~ (p103(X17) & ~ p104(X17) & p4(X17)) | ~ r1(X0, X17)) & ~ ! [X18] : (~ (p103(X18) & ~ p104(X18) & ~ p4(X18)) | ~ r1(X0, X18)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X19] : (~ (p104(X19) & ~ p105(X19) & p5(X19)) | ~ r1(X0, X19)) & ~ ! [X20] : (~ (p104(X20) & ~ p105(X20) & ~ p5(X20)) | ~ r1(X0, X20)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X21] : (~ (p105(X21) & ~ p106(X21) & p6(X21)) | ~ r1(X0, X21)) & ~ ! [X22] : (~ (p105(X22) & ~ p106(X22) & ~ p6(X22)) | ~ r1(X0, X22)))) & ! [X23] : (((~ p101(X23) | p100(X23)) & (~ p102(X23) | p101(X23)) & (~ p103(X23) | p102(X23)) & (~ p104(X23) | p103(X23)) & (~ p105(X23) | p104(X23)) & (~ p106(X23) | p105(X23)) & (~ p100(X23) | ((~ p1(X23) | ! [X24] : (~ p100(X24) | p1(X24) | ~ r1(X23, X24))) & (p1(X23) | ! [X25] : (~ p100(X25) | ~ p1(X25) | ~ r1(X23, X25))))) & (~ p101(X23) | ((~ p2(X23) | ! [X26] : (~ p101(X26) | p2(X26) | ~ r1(X23, X26))) & (p2(X23) | ! [X27] : (~ p101(X27) | ~ p2(X27) | ~ r1(X23, X27))))) & (~ p102(X23) | ((~ p3(X23) | ! [X28] : (~ p102(X28) | p3(X28) | ~ r1(X23, X28))) & (p3(X23) | ! [X29] : (~ p102(X29) | ~ p3(X29) | ~ r1(X23, X29))))) & (~ p103(X23) | ((~ p4(X23) | ! [X30] : (~ p103(X30) | p4(X30) | ~ r1(X23, X30))) & (p4(X23) | ! [X31] : (~ p103(X31) | ~ p4(X31) | ~ r1(X23, X31))))) & (~ p104(X23) | ((~ p5(X23) | ! [X32] : (~ p104(X32) | p5(X32) | ~ r1(X23, X32))) & (p5(X23) | ! [X33] : (~ p104(X33) | ~ p5(X33) | ~ r1(X23, X33))))) & (~ p105(X23) | ((~ p6(X23) | ! [X34] : (~ p105(X34) | p6(X34) | ~ r1(X23, X34))) & (p6(X23) | ! [X35] : (~ p105(X35) | ~ p6(X35) | ~ r1(X23, X35))))) & (~ (p100(X23) & ~ p101(X23)) | (~ ! [X36] : (~ (p101(X36) & ~ p102(X36) & p2(X36)) | ~ r1(X23, X36)) & ~ ! [X37] : (~ (p101(X37) & ~ p102(X37) & ~ p2(X37)) | ~ r1(X23, X37)))) & (~ (p101(X23) & ~ p102(X23)) | (~ ! [X38] : (~ (p102(X38) & ~ p103(X38) & p3(X38)) | ~ r1(X23, X38)) & ~ ! [X39] : (~ (p102(X39) & ~ p103(X39) & ~ p3(X39)) | ~ r1(X23, X39)))) & (~ (p102(X23) & ~ p103(X23)) | (~ ! [X40] : (~ (p103(X40) & ~ p104(X40) & p4(X40)) | ~ r1(X23, X40)) & ~ ! [X41] : (~ (p103(X41) & ~ p104(X41) & ~ p4(X41)) | ~ r1(X23, X41)))) & (~ (p103(X23) & ~ p104(X23)) | (~ ! [X42] : (~ (p104(X42) & ~ p105(X42) & p5(X42)) | ~ r1(X23, X42)) & ~ ! [X43] : (~ (p104(X43) & ~ p105(X43) & ~ p5(X43)) | ~ r1(X23, X43)))) & (~ (p104(X23) & ~ p105(X23)) | (~ ! [X44] : (~ (p105(X44) & ~ p106(X44) & p6(X44)) | ~ r1(X23, X44)) & ~ ! [X45] : (~ (p105(X45) & ~ p106(X45) & ~ p6(X45)) | ~ r1(X23, X45))))) | ~ r1(X0, X23)) & ! [X46] : (! [X47] : (((~ p101(X47) | p100(X47)) & (~ p102(X47) | p101(X47)) & (~ p103(X47) | p102(X47)) & (~ p104(X47) | p103(X47)) & (~ p105(X47) | p104(X47)) & (~ p106(X47) | p105(X47)) & (~ p100(X47) | ((~ p1(X47) | ! [X48] : (~ p100(X48) | p1(X48) | ~ r1(X47, X48))) & (p1(X47) | ! [X49] : (~ p100(X49) | ~ p1(X49) | ~ r1(X47, X49))))) & (~ p101(X47) | ((~ p2(X47) | ! [X50] : (~ p101(X50) | p2(X50) | ~ r1(X47, X50))) & (p2(X47) | ! [X51] : (~ p101(X51) | ~ p2(X51) | ~ r1(X47, X51))))) & (~ p102(X47) | ((~ p3(X47) | ! [X52] : (~ p102(X52) | p3(X52) | ~ r1(X47, X52))) & (p3(X47) | ! [X53] : (~ p102(X53) | ~ p3(X53) | ~ r1(X47, X53))))) & (~ p103(X47) | ((~ p4(X47) | ! [X54] : (~ p103(X54) | p4(X54) | ~ r1(X47, X54))) & (p4(X47) | ! [X55] : (~ p103(X55) | ~ p4(X55) | ~ r1(X47, X55))))) & (~ p104(X47) | ((~ p5(X47) | ! [X56] : (~ p104(X56) | p5(X56) | ~ r1(X47, X56))) & (p5(X47) | ! [X57] : (~ p104(X57) | ~ p5(X57) | ~ r1(X47, X57))))) & (~ p105(X47) | ((~ p6(X47) | ! [X58] : (~ p105(X58) | p6(X58) | ~ r1(X47, X58))) & (p6(X47) | ! [X59] : (~ p105(X59) | ~ p6(X59) | ~ r1(X47, X59))))) & (~ (p100(X47) & ~ p101(X47)) | (~ ! [X60] : (~ (p101(X60) & ~ p102(X60) & p2(X60)) | ~ r1(X47, X60)) & ~ ! [X61] : (~ (p101(X61) & ~ p102(X61) & ~ p2(X61)) | ~ r1(X47, X61)))) & (~ (p101(X47) & ~ p102(X47)) | (~ ! [X62] : (~ (p102(X62) & ~ p103(X62) & p3(X62)) | ~ r1(X47, X62)) & ~ ! [X63] : (~ (p102(X63) & ~ p103(X63) & ~ p3(X63)) | ~ r1(X47, X63)))) & (~ (p102(X47) & ~ p103(X47)) | (~ ! [X64] : (~ (p103(X64) & ~ p104(X64) & p4(X64)) | ~ r1(X47, X64)) & ~ ! [X65] : (~ (p103(X65) & ~ p104(X65) & ~ p4(X65)) | ~ r1(X47, X65)))) & (~ (p103(X47) & ~ p104(X47)) | (~ ! [X66] : (~ (p104(X66) & ~ p105(X66) & p5(X66)) | ~ r1(X47, X66)) & ~ ! [X67] : (~ (p104(X67) & ~ p105(X67) & ~ p5(X67)) | ~ r1(X47, X67)))) & (~ (p104(X47) & ~ p105(X47)) | (~ ! [X68] : (~ (p105(X68) & ~ p106(X68) & p6(X68)) | ~ r1(X47, X68)) & ~ ! [X69] : (~ (p105(X69) & ~ p106(X69) & ~ p6(X69)) | ~ r1(X47, X69))))) | ~ r1(X46, X47)) | ~ r1(X0, X46)) & ! [X70] : (! [X71] : (! [X72] : (((~ p101(X72) | p100(X72)) & (~ p102(X72) | p101(X72)) & (~ p103(X72) | p102(X72)) & (~ p104(X72) | p103(X72)) & (~ p105(X72) | p104(X72)) & (~ p106(X72) | p105(X72)) & (~ p100(X72) | ((~ p1(X72) | ! [X73] : (~ p100(X73) | p1(X73) | ~ r1(X72, X73))) & (p1(X72) | ! [X74] : (~ p100(X74) | ~ p1(X74) | ~ r1(X72, X74))))) & (~ p101(X72) | ((~ p2(X72) | ! [X75] : (~ p101(X75) | p2(X75) | ~ r1(X72, X75))) & (p2(X72) | ! [X76] : (~ p101(X76) | ~ p2(X76) | ~ r1(X72, X76))))) & (~ p102(X72) | ((~ p3(X72) | ! [X77] : (~ p102(X77) | p3(X77) | ~ r1(X72, X77))) & (p3(X72) | ! [X78] : (~ p102(X78) | ~ p3(X78) | ~ r1(X72, X78))))) & (~ p103(X72) | ((~ p4(X72) | ! [X79] : (~ p103(X79) | p4(X79) | ~ r1(X72, X79))) & (p4(X72) | ! [X80] : (~ p103(X80) | ~ p4(X80) | ~ r1(X72, X80))))) & (~ p104(X72) | ((~ p5(X72) | ! [X81] : (~ p104(X81) | p5(X81) | ~ r1(X72, X81))) & (p5(X72) | ! [X82] : (~ p104(X82) | ~ p5(X82) | ~ r1(X72, X82))))) & (~ p105(X72) | ((~ p6(X72) | ! [X83] : (~ p105(X83) | p6(X83) | ~ r1(X72, X83))) & (p6(X72) | ! [X84] : (~ p105(X84) | ~ p6(X84) | ~ r1(X72, X84))))) & (~ (p100(X72) & ~ p101(X72)) | (~ ! [X85] : (~ (p101(X85) & ~ p102(X85) & p2(X85)) | ~ r1(X72, X85)) & ~ ! [X86] : (~ (p101(X86) & ~ p102(X86) & ~ p2(X86)) | ~ r1(X72, X86)))) & (~ (p101(X72) & ~ p102(X72)) | (~ ! [X87] : (~ (p102(X87) & ~ p103(X87) & p3(X87)) | ~ r1(X72, X87)) & ~ ! [X88] : (~ (p102(X88) & ~ p103(X88) & ~ p3(X88)) | ~ r1(X72, X88)))) & (~ (p102(X72) & ~ p103(X72)) | (~ ! [X89] : (~ (p103(X89) & ~ p104(X89) & p4(X89)) | ~ r1(X72, X89)) & ~ ! [X90] : (~ (p103(X90) & ~ p104(X90) & ~ p4(X90)) | ~ r1(X72, X90)))) & (~ (p103(X72) & ~ p104(X72)) | (~ ! [X91] : (~ (p104(X91) & ~ p105(X91) & p5(X91)) | ~ r1(X72, X91)) & ~ ! [X92] : (~ (p104(X92) & ~ p105(X92) & ~ p5(X92)) | ~ r1(X72, X92)))) & (~ (p104(X72) & ~ p105(X72)) | (~ ! [X93] : (~ (p105(X93) & ~ p106(X93) & p6(X93)) | ~ r1(X72, X93)) & ~ ! [X94] : (~ (p105(X94) & ~ p106(X94) & ~ p6(X94)) | ~ r1(X72, X94))))) | ~ r1(X71, X72)) | ~ r1(X70, X71)) | ~ r1(X0, X70)) & ! [X95] : (! [X96] : (! [X97] : (! [X98] : (((~ p101(X98) | p100(X98)) & (~ p102(X98) | p101(X98)) & (~ p103(X98) | p102(X98)) & (~ p104(X98) | p103(X98)) & (~ p105(X98) | p104(X98)) & (~ p106(X98) | p105(X98)) & (~ p100(X98) | ((~ p1(X98) | ! [X99] : (~ p100(X99) | p1(X99) | ~ r1(X98, X99))) & (p1(X98) | ! [X100] : (~ p100(X100) | ~ p1(X100) | ~ r1(X98, X100))))) & (~ p101(X98) | ((~ p2(X98) | ! [X101] : (~ p101(X101) | p2(X101) | ~ r1(X98, X101))) & (p2(X98) | ! [X102] : (~ p101(X102) | ~ p2(X102) | ~ r1(X98, X102))))) & (~ p102(X98) | ((~ p3(X98) | ! [X103] : (~ p102(X103) | p3(X103) | ~ r1(X98, X103))) & (p3(X98) | ! [X104] : (~ p102(X104) | ~ p3(X104) | ~ r1(X98, X104))))) & (~ p103(X98) | ((~ p4(X98) | ! [X105] : (~ p103(X105) | p4(X105) | ~ r1(X98, X105))) & (p4(X98) | ! [X106] : (~ p103(X106) | ~ p4(X106) | ~ r1(X98, X106))))) & (~ p104(X98) | ((~ p5(X98) | ! [X107] : (~ p104(X107) | p5(X107) | ~ r1(X98, X107))) & (p5(X98) | ! [X108] : (~ p104(X108) | ~ p5(X108) | ~ r1(X98, X108))))) & (~ p105(X98) | ((~ p6(X98) | ! [X109] : (~ p105(X109) | p6(X109) | ~ r1(X98, X109))) & (p6(X98) | ! [X110] : (~ p105(X110) | ~ p6(X110) | ~ r1(X98, X110))))) & (~ (p100(X98) & ~ p101(X98)) | (~ ! [X111] : (~ (p101(X111) & ~ p102(X111) & p2(X111)) | ~ r1(X98, X111)) & ~ ! [X112] : (~ (p101(X112) & ~ p102(X112) & ~ p2(X112)) | ~ r1(X98, X112)))) & (~ (p101(X98) & ~ p102(X98)) | (~ ! [X113] : (~ (p102(X113) & ~ p103(X113) & p3(X113)) | ~ r1(X98, X113)) & ~ ! [X114] : (~ (p102(X114) & ~ p103(X114) & ~ p3(X114)) | ~ r1(X98, X114)))) & (~ (p102(X98) & ~ p103(X98)) | (~ ! [X115] : (~ (p103(X115) & ~ p104(X115) & p4(X115)) | ~ r1(X98, X115)) & ~ ! [X116] : (~ (p103(X116) & ~ p104(X116) & ~ p4(X116)) | ~ r1(X98, X116)))) & (~ (p103(X98) & ~ p104(X98)) | (~ ! [X117] : (~ (p104(X117) & ~ p105(X117) & p5(X117)) | ~ r1(X98, X117)) & ~ ! [X118] : (~ (p104(X118) & ~ p105(X118) & ~ p5(X118)) | ~ r1(X98, X118)))) & (~ (p104(X98) & ~ p105(X98)) | (~ ! [X119] : (~ (p105(X119) & ~ p106(X119) & p6(X119)) | ~ r1(X98, X119)) & ~ ! [X120] : (~ (p105(X120) & ~ p106(X120) & ~ p6(X120)) | ~ r1(X98, X120))))) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X0, X95)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (((~ p101(X125) | p100(X125)) & (~ p102(X125) | p101(X125)) & (~ p103(X125) | p102(X125)) & (~ p104(X125) | p103(X125)) & (~ p105(X125) | p104(X125)) & (~ p106(X125) | p105(X125)) & (~ p100(X125) | ((~ p1(X125) | ! [X126] : (~ p100(X126) | p1(X126) | ~ r1(X125, X126))) & (p1(X125) | ! [X127] : (~ p100(X127) | ~ p1(X127) | ~ r1(X125, X127))))) & (~ p101(X125) | ((~ p2(X125) | ! [X128] : (~ p101(X128) | p2(X128) | ~ r1(X125, X128))) & (p2(X125) | ! [X129] : (~ p101(X129) | ~ p2(X129) | ~ r1(X125, X129))))) & (~ p102(X125) | ((~ p3(X125) | ! [X130] : (~ p102(X130) | p3(X130) | ~ r1(X125, X130))) & (p3(X125) | ! [X131] : (~ p102(X131) | ~ p3(X131) | ~ r1(X125, X131))))) & (~ p103(X125) | ((~ p4(X125) | ! [X132] : (~ p103(X132) | p4(X132) | ~ r1(X125, X132))) & (p4(X125) | ! [X133] : (~ p103(X133) | ~ p4(X133) | ~ r1(X125, X133))))) & (~ p104(X125) | ((~ p5(X125) | ! [X134] : (~ p104(X134) | p5(X134) | ~ r1(X125, X134))) & (p5(X125) | ! [X135] : (~ p104(X135) | ~ p5(X135) | ~ r1(X125, X135))))) & (~ p105(X125) | ((~ p6(X125) | ! [X136] : (~ p105(X136) | p6(X136) | ~ r1(X125, X136))) & (p6(X125) | ! [X137] : (~ p105(X137) | ~ p6(X137) | ~ r1(X125, X137))))) & (~ (p100(X125) & ~ p101(X125)) | (~ ! [X138] : (~ (p101(X138) & ~ p102(X138) & p2(X138)) | ~ r1(X125, X138)) & ~ ! [X139] : (~ (p101(X139) & ~ p102(X139) & ~ p2(X139)) | ~ r1(X125, X139)))) & (~ (p101(X125) & ~ p102(X125)) | (~ ! [X140] : (~ (p102(X140) & ~ p103(X140) & p3(X140)) | ~ r1(X125, X140)) & ~ ! [X141] : (~ (p102(X141) & ~ p103(X141) & ~ p3(X141)) | ~ r1(X125, X141)))) & (~ (p102(X125) & ~ p103(X125)) | (~ ! [X142] : (~ (p103(X142) & ~ p104(X142) & p4(X142)) | ~ r1(X125, X142)) & ~ ! [X143] : (~ (p103(X143) & ~ p104(X143) & ~ p4(X143)) | ~ r1(X125, X143)))) & (~ (p103(X125) & ~ p104(X125)) | (~ ! [X144] : (~ (p104(X144) & ~ p105(X144) & p5(X144)) | ~ r1(X125, X144)) & ~ ! [X145] : (~ (p104(X145) & ~ p105(X145) & ~ p5(X145)) | ~ r1(X125, X145)))) & (~ (p104(X125) & ~ p105(X125)) | (~ ! [X146] : (~ (p105(X146) & ~ p106(X146) & p6(X146)) | ~ r1(X125, X146)) & ~ ! [X147] : (~ (p105(X147) & ~ p106(X147) & ~ p6(X147)) | ~ r1(X125, X147))))) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121))) | ~ ! [X148] : (! [X149] : (! [X150] : (! [X151] : (! [X152] : (p3(X152) | ~ r1(X151, X152)) | ~ r1(X150, X151)) | ~ r1(X149, X150)) | ~ r1(X148, X149)) | ~ r1(X0, X148))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X2] : (~ p100(X2) | ~ p1(X2) | ~ r1(X0, X2))))) & (~ p101(X0) | ((~ p2(X0) | ! [X3] : (~ p101(X3) | p2(X3) | ~ r1(X0, X3))) & (p2(X0) | ! [X4] : (~ p101(X4) | ~ p2(X4) | ~ r1(X0, X4))))) & (~ p102(X0) | ((~ p3(X0) | ! [X5] : (~ p102(X5) | p3(X5) | ~ r1(X0, X5))) & (p3(X0) | ! [X6] : (~ p102(X6) | ~ p3(X6) | ~ r1(X0, X6))))) & (~ p103(X0) | ((~ p4(X0) | ! [X7] : (~ p103(X7) | p4(X7) | ~ r1(X0, X7))) & (p4(X0) | ! [X8] : (~ p103(X8) | ~ p4(X8) | ~ r1(X0, X8))))) & (~ p104(X0) | ((~ p5(X0) | ! [X9] : (~ p104(X9) | p5(X9) | ~ r1(X0, X9))) & (p5(X0) | ! [X10] : (~ p104(X10) | ~ p5(X10) | ~ r1(X0, X10))))) & (~ p105(X0) | ((~ p6(X0) | ! [X11] : (~ p105(X11) | p6(X11) | ~ r1(X0, X11))) & (p6(X0) | ! [X12] : (~ p105(X12) | ~ p6(X12) | ~ r1(X0, X12))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X13] : (~ (p101(X13) & ~ p102(X13) & p2(X13)) | ~ r1(X0, X13)) & ~ ! [X14] : (~ (p101(X14) & ~ p102(X14) & ~ p2(X14)) | ~ r1(X0, X14)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X15] : (~ (p102(X15) & ~ p103(X15) & p3(X15)) | ~ r1(X0, X15)) & ~ ! [X16] : (~ (p102(X16) & ~ p103(X16) & ~ p3(X16)) | ~ r1(X0, X16)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X17] : (~ (p103(X17) & ~ p104(X17) & p4(X17)) | ~ r1(X0, X17)) & ~ ! [X18] : (~ (p103(X18) & ~ p104(X18) & ~ p4(X18)) | ~ r1(X0, X18)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X19] : (~ (p104(X19) & ~ p105(X19) & p5(X19)) | ~ r1(X0, X19)) & ~ ! [X20] : (~ (p104(X20) & ~ p105(X20) & ~ p5(X20)) | ~ r1(X0, X20)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X21] : (~ (p105(X21) & ~ p106(X21) & p6(X21)) | ~ r1(X0, X21)) & ~ ! [X22] : (~ (p105(X22) & ~ p106(X22) & ~ p6(X22)) | ~ r1(X0, X22)))) & ! [X23] : (((~ p101(X23) | p100(X23)) & (~ p102(X23) | p101(X23)) & (~ p103(X23) | p102(X23)) & (~ p104(X23) | p103(X23)) & (~ p105(X23) | p104(X23)) & (~ p106(X23) | p105(X23)) & (~ p100(X23) | ((~ p1(X23) | ! [X24] : (~ p100(X24) | p1(X24) | ~ r1(X23, X24))) & (p1(X23) | ! [X25] : (~ p100(X25) | ~ p1(X25) | ~ r1(X23, X25))))) & (~ p101(X23) | ((~ p2(X23) | ! [X26] : (~ p101(X26) | p2(X26) | ~ r1(X23, X26))) & (p2(X23) | ! [X27] : (~ p101(X27) | ~ p2(X27) | ~ r1(X23, X27))))) & (~ p102(X23) | ((~ p3(X23) | ! [X28] : (~ p102(X28) | p3(X28) | ~ r1(X23, X28))) & (p3(X23) | ! [X29] : (~ p102(X29) | ~ p3(X29) | ~ r1(X23, X29))))) & (~ p103(X23) | ((~ p4(X23) | ! [X30] : (~ p103(X30) | p4(X30) | ~ r1(X23, X30))) & (p4(X23) | ! [X31] : (~ p103(X31) | ~ p4(X31) | ~ r1(X23, X31))))) & (~ p104(X23) | ((~ p5(X23) | ! [X32] : (~ p104(X32) | p5(X32) | ~ r1(X23, X32))) & (p5(X23) | ! [X33] : (~ p104(X33) | ~ p5(X33) | ~ r1(X23, X33))))) & (~ p105(X23) | ((~ p6(X23) | ! [X34] : (~ p105(X34) | p6(X34) | ~ r1(X23, X34))) & (p6(X23) | ! [X35] : (~ p105(X35) | ~ p6(X35) | ~ r1(X23, X35))))) & (~ (p100(X23) & ~ p101(X23)) | (~ ! [X36] : (~ (p101(X36) & ~ p102(X36) & p2(X36)) | ~ r1(X23, X36)) & ~ ! [X37] : (~ (p101(X37) & ~ p102(X37) & ~ p2(X37)) | ~ r1(X23, X37)))) & (~ (p101(X23) & ~ p102(X23)) | (~ ! [X38] : (~ (p102(X38) & ~ p103(X38) & p3(X38)) | ~ r1(X23, X38)) & ~ ! [X39] : (~ (p102(X39) & ~ p103(X39) & ~ p3(X39)) | ~ r1(X23, X39)))) & (~ (p102(X23) & ~ p103(X23)) | (~ ! [X40] : (~ (p103(X40) & ~ p104(X40) & p4(X40)) | ~ r1(X23, X40)) & ~ ! [X41] : (~ (p103(X41) & ~ p104(X41) & ~ p4(X41)) | ~ r1(X23, X41)))) & (~ (p103(X23) & ~ p104(X23)) | (~ ! [X42] : (~ (p104(X42) & ~ p105(X42) & p5(X42)) | ~ r1(X23, X42)) & ~ ! [X43] : (~ (p104(X43) & ~ p105(X43) & ~ p5(X43)) | ~ r1(X23, X43)))) & (~ (p104(X23) & ~ p105(X23)) | (~ ! [X44] : (~ (p105(X44) & ~ p106(X44) & p6(X44)) | ~ r1(X23, X44)) & ~ ! [X45] : (~ (p105(X45) & ~ p106(X45) & ~ p6(X45)) | ~ r1(X23, X45))))) | ~ r1(X0, X23)) & ! [X46] : (! [X47] : (((~ p101(X47) | p100(X47)) & (~ p102(X47) | p101(X47)) & (~ p103(X47) | p102(X47)) & (~ p104(X47) | p103(X47)) & (~ p105(X47) | p104(X47)) & (~ p106(X47) | p105(X47)) & (~ p100(X47) | ((~ p1(X47) | ! [X48] : (~ p100(X48) | p1(X48) | ~ r1(X47, X48))) & (p1(X47) | ! [X49] : (~ p100(X49) | ~ p1(X49) | ~ r1(X47, X49))))) & (~ p101(X47) | ((~ p2(X47) | ! [X50] : (~ p101(X50) | p2(X50) | ~ r1(X47, X50))) & (p2(X47) | ! [X51] : (~ p101(X51) | ~ p2(X51) | ~ r1(X47, X51))))) & (~ p102(X47) | ((~ p3(X47) | ! [X52] : (~ p102(X52) | p3(X52) | ~ r1(X47, X52))) & (p3(X47) | ! [X53] : (~ p102(X53) | ~ p3(X53) | ~ r1(X47, X53))))) & (~ p103(X47) | ((~ p4(X47) | ! [X54] : (~ p103(X54) | p4(X54) | ~ r1(X47, X54))) & (p4(X47) | ! [X55] : (~ p103(X55) | ~ p4(X55) | ~ r1(X47, X55))))) & (~ p104(X47) | ((~ p5(X47) | ! [X56] : (~ p104(X56) | p5(X56) | ~ r1(X47, X56))) & (p5(X47) | ! [X57] : (~ p104(X57) | ~ p5(X57) | ~ r1(X47, X57))))) & (~ p105(X47) | ((~ p6(X47) | ! [X58] : (~ p105(X58) | p6(X58) | ~ r1(X47, X58))) & (p6(X47) | ! [X59] : (~ p105(X59) | ~ p6(X59) | ~ r1(X47, X59))))) & (~ (p100(X47) & ~ p101(X47)) | (~ ! [X60] : (~ (p101(X60) & ~ p102(X60) & p2(X60)) | ~ r1(X47, X60)) & ~ ! [X61] : (~ (p101(X61) & ~ p102(X61) & ~ p2(X61)) | ~ r1(X47, X61)))) & (~ (p101(X47) & ~ p102(X47)) | (~ ! [X62] : (~ (p102(X62) & ~ p103(X62) & p3(X62)) | ~ r1(X47, X62)) & ~ ! [X63] : (~ (p102(X63) & ~ p103(X63) & ~ p3(X63)) | ~ r1(X47, X63)))) & (~ (p102(X47) & ~ p103(X47)) | (~ ! [X64] : (~ (p103(X64) & ~ p104(X64) & p4(X64)) | ~ r1(X47, X64)) & ~ ! [X65] : (~ (p103(X65) & ~ p104(X65) & ~ p4(X65)) | ~ r1(X47, X65)))) & (~ (p103(X47) & ~ p104(X47)) | (~ ! [X66] : (~ (p104(X66) & ~ p105(X66) & p5(X66)) | ~ r1(X47, X66)) & ~ ! [X67] : (~ (p104(X67) & ~ p105(X67) & ~ p5(X67)) | ~ r1(X47, X67)))) & (~ (p104(X47) & ~ p105(X47)) | (~ ! [X68] : (~ (p105(X68) & ~ p106(X68) & p6(X68)) | ~ r1(X47, X68)) & ~ ! [X69] : (~ (p105(X69) & ~ p106(X69) & ~ p6(X69)) | ~ r1(X47, X69))))) | ~ r1(X46, X47)) | ~ r1(X0, X46)) & ! [X70] : (! [X71] : (! [X72] : (((~ p101(X72) | p100(X72)) & (~ p102(X72) | p101(X72)) & (~ p103(X72) | p102(X72)) & (~ p104(X72) | p103(X72)) & (~ p105(X72) | p104(X72)) & (~ p106(X72) | p105(X72)) & (~ p100(X72) | ((~ p1(X72) | ! [X73] : (~ p100(X73) | p1(X73) | ~ r1(X72, X73))) & (p1(X72) | ! [X74] : (~ p100(X74) | ~ p1(X74) | ~ r1(X72, X74))))) & (~ p101(X72) | ((~ p2(X72) | ! [X75] : (~ p101(X75) | p2(X75) | ~ r1(X72, X75))) & (p2(X72) | ! [X76] : (~ p101(X76) | ~ p2(X76) | ~ r1(X72, X76))))) & (~ p102(X72) | ((~ p3(X72) | ! [X77] : (~ p102(X77) | p3(X77) | ~ r1(X72, X77))) & (p3(X72) | ! [X78] : (~ p102(X78) | ~ p3(X78) | ~ r1(X72, X78))))) & (~ p103(X72) | ((~ p4(X72) | ! [X79] : (~ p103(X79) | p4(X79) | ~ r1(X72, X79))) & (p4(X72) | ! [X80] : (~ p103(X80) | ~ p4(X80) | ~ r1(X72, X80))))) & (~ p104(X72) | ((~ p5(X72) | ! [X81] : (~ p104(X81) | p5(X81) | ~ r1(X72, X81))) & (p5(X72) | ! [X82] : (~ p104(X82) | ~ p5(X82) | ~ r1(X72, X82))))) & (~ p105(X72) | ((~ p6(X72) | ! [X83] : (~ p105(X83) | p6(X83) | ~ r1(X72, X83))) & (p6(X72) | ! [X84] : (~ p105(X84) | ~ p6(X84) | ~ r1(X72, X84))))) & (~ (p100(X72) & ~ p101(X72)) | (~ ! [X85] : (~ (p101(X85) & ~ p102(X85) & p2(X85)) | ~ r1(X72, X85)) & ~ ! [X86] : (~ (p101(X86) & ~ p102(X86) & ~ p2(X86)) | ~ r1(X72, X86)))) & (~ (p101(X72) & ~ p102(X72)) | (~ ! [X87] : (~ (p102(X87) & ~ p103(X87) & p3(X87)) | ~ r1(X72, X87)) & ~ ! [X88] : (~ (p102(X88) & ~ p103(X88) & ~ p3(X88)) | ~ r1(X72, X88)))) & (~ (p102(X72) & ~ p103(X72)) | (~ ! [X89] : (~ (p103(X89) & ~ p104(X89) & p4(X89)) | ~ r1(X72, X89)) & ~ ! [X90] : (~ (p103(X90) & ~ p104(X90) & ~ p4(X90)) | ~ r1(X72, X90)))) & (~ (p103(X72) & ~ p104(X72)) | (~ ! [X91] : (~ (p104(X91) & ~ p105(X91) & p5(X91)) | ~ r1(X72, X91)) & ~ ! [X92] : (~ (p104(X92) & ~ p105(X92) & ~ p5(X92)) | ~ r1(X72, X92)))) & (~ (p104(X72) & ~ p105(X72)) | (~ ! [X93] : (~ (p105(X93) & ~ p106(X93) & p6(X93)) | ~ r1(X72, X93)) & ~ ! [X94] : (~ (p105(X94) & ~ p106(X94) & ~ p6(X94)) | ~ r1(X72, X94))))) | ~ r1(X71, X72)) | ~ r1(X70, X71)) | ~ r1(X0, X70)) & ! [X95] : (! [X96] : (! [X97] : (! [X98] : (((~ p101(X98) | p100(X98)) & (~ p102(X98) | p101(X98)) & (~ p103(X98) | p102(X98)) & (~ p104(X98) | p103(X98)) & (~ p105(X98) | p104(X98)) & (~ p106(X98) | p105(X98)) & (~ p100(X98) | ((~ p1(X98) | ! [X99] : (~ p100(X99) | p1(X99) | ~ r1(X98, X99))) & (p1(X98) | ! [X100] : (~ p100(X100) | ~ p1(X100) | ~ r1(X98, X100))))) & (~ p101(X98) | ((~ p2(X98) | ! [X101] : (~ p101(X101) | p2(X101) | ~ r1(X98, X101))) & (p2(X98) | ! [X102] : (~ p101(X102) | ~ p2(X102) | ~ r1(X98, X102))))) & (~ p102(X98) | ((~ p3(X98) | ! [X103] : (~ p102(X103) | p3(X103) | ~ r1(X98, X103))) & (p3(X98) | ! [X104] : (~ p102(X104) | ~ p3(X104) | ~ r1(X98, X104))))) & (~ p103(X98) | ((~ p4(X98) | ! [X105] : (~ p103(X105) | p4(X105) | ~ r1(X98, X105))) & (p4(X98) | ! [X106] : (~ p103(X106) | ~ p4(X106) | ~ r1(X98, X106))))) & (~ p104(X98) | ((~ p5(X98) | ! [X107] : (~ p104(X107) | p5(X107) | ~ r1(X98, X107))) & (p5(X98) | ! [X108] : (~ p104(X108) | ~ p5(X108) | ~ r1(X98, X108))))) & (~ p105(X98) | ((~ p6(X98) | ! [X109] : (~ p105(X109) | p6(X109) | ~ r1(X98, X109))) & (p6(X98) | ! [X110] : (~ p105(X110) | ~ p6(X110) | ~ r1(X98, X110))))) & (~ (p100(X98) & ~ p101(X98)) | (~ ! [X111] : (~ (p101(X111) & ~ p102(X111) & p2(X111)) | ~ r1(X98, X111)) & ~ ! [X112] : (~ (p101(X112) & ~ p102(X112) & ~ p2(X112)) | ~ r1(X98, X112)))) & (~ (p101(X98) & ~ p102(X98)) | (~ ! [X113] : (~ (p102(X113) & ~ p103(X113) & p3(X113)) | ~ r1(X98, X113)) & ~ ! [X114] : (~ (p102(X114) & ~ p103(X114) & ~ p3(X114)) | ~ r1(X98, X114)))) & (~ (p102(X98) & ~ p103(X98)) | (~ ! [X115] : (~ (p103(X115) & ~ p104(X115) & p4(X115)) | ~ r1(X98, X115)) & ~ ! [X116] : (~ (p103(X116) & ~ p104(X116) & ~ p4(X116)) | ~ r1(X98, X116)))) & (~ (p103(X98) & ~ p104(X98)) | (~ ! [X117] : (~ (p104(X117) & ~ p105(X117) & p5(X117)) | ~ r1(X98, X117)) & ~ ! [X118] : (~ (p104(X118) & ~ p105(X118) & ~ p5(X118)) | ~ r1(X98, X118)))) & (~ (p104(X98) & ~ p105(X98)) | (~ ! [X119] : (~ (p105(X119) & ~ p106(X119) & p6(X119)) | ~ r1(X98, X119)) & ~ ! [X120] : (~ (p105(X120) & ~ p106(X120) & ~ p6(X120)) | ~ r1(X98, X120))))) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X0, X95)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (((~ p101(X125) | p100(X125)) & (~ p102(X125) | p101(X125)) & (~ p103(X125) | p102(X125)) & (~ p104(X125) | p103(X125)) & (~ p105(X125) | p104(X125)) & (~ p106(X125) | p105(X125)) & (~ p100(X125) | ((~ p1(X125) | ! [X126] : (~ p100(X126) | p1(X126) | ~ r1(X125, X126))) & (p1(X125) | ! [X127] : (~ p100(X127) | ~ p1(X127) | ~ r1(X125, X127))))) & (~ p101(X125) | ((~ p2(X125) | ! [X128] : (~ p101(X128) | p2(X128) | ~ r1(X125, X128))) & (p2(X125) | ! [X129] : (~ p101(X129) | ~ p2(X129) | ~ r1(X125, X129))))) & (~ p102(X125) | ((~ p3(X125) | ! [X130] : (~ p102(X130) | p3(X130) | ~ r1(X125, X130))) & (p3(X125) | ! [X131] : (~ p102(X131) | ~ p3(X131) | ~ r1(X125, X131))))) & (~ p103(X125) | ((~ p4(X125) | ! [X132] : (~ p103(X132) | p4(X132) | ~ r1(X125, X132))) & (p4(X125) | ! [X133] : (~ p103(X133) | ~ p4(X133) | ~ r1(X125, X133))))) & (~ p104(X125) | ((~ p5(X125) | ! [X134] : (~ p104(X134) | p5(X134) | ~ r1(X125, X134))) & (p5(X125) | ! [X135] : (~ p104(X135) | ~ p5(X135) | ~ r1(X125, X135))))) & (~ p105(X125) | ((~ p6(X125) | ! [X136] : (~ p105(X136) | p6(X136) | ~ r1(X125, X136))) & (p6(X125) | ! [X137] : (~ p105(X137) | ~ p6(X137) | ~ r1(X125, X137))))) & (~ (p100(X125) & ~ p101(X125)) | (~ ! [X138] : (~ (p101(X138) & ~ p102(X138) & p2(X138)) | ~ r1(X125, X138)) & ~ ! [X139] : (~ (p101(X139) & ~ p102(X139) & ~ p2(X139)) | ~ r1(X125, X139)))) & (~ (p101(X125) & ~ p102(X125)) | (~ ! [X140] : (~ (p102(X140) & ~ p103(X140) & p3(X140)) | ~ r1(X125, X140)) & ~ ! [X141] : (~ (p102(X141) & ~ p103(X141) & ~ p3(X141)) | ~ r1(X125, X141)))) & (~ (p102(X125) & ~ p103(X125)) | (~ ! [X142] : (~ (p103(X142) & ~ p104(X142) & p4(X142)) | ~ r1(X125, X142)) & ~ ! [X143] : (~ (p103(X143) & ~ p104(X143) & ~ p4(X143)) | ~ r1(X125, X143)))) & (~ (p103(X125) & ~ p104(X125)) | (~ ! [X144] : (~ (p104(X144) & ~ p105(X144) & p5(X144)) | ~ r1(X125, X144)) & ~ ! [X145] : (~ (p104(X145) & ~ p105(X145) & ~ p5(X145)) | ~ r1(X125, X145)))) & (~ (p104(X125) & ~ p105(X125)) | (~ ! [X146] : (~ (p105(X146) & ~ p106(X146) & p6(X146)) | ~ r1(X125, X146)) & ~ ! [X147] : (~ (p105(X147) & ~ p106(X147) & ~ p6(X147)) | ~ r1(X125, X147))))) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121))) | ~ ! [X148] : (! [X149] : (! [X150] : (! [X151] : (! [X152] : (p3(X152) | ~ r1(X151, X152)) | ~ r1(X150, X151)) | ~ r1(X149, X150)) | ~ r1(X148, X149)) | ~ r1(X0, X148))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X1] : (~ p100(X1) | ~ p1(X1) | ~ r1(X0, X1))))) & (~ p101(X0) | ((~ p2(X0) | ! [X1] : (~ p101(X1) | p2(X1) | ~ r1(X0, X1))) & (p2(X0) | ! [X1] : (~ p101(X1) | ~ p2(X1) | ~ r1(X0, X1))))) & (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X1] : (~ p102(X1) | ~ p3(X1) | ~ r1(X0, X1))))) & (~ p103(X0) | ((~ p4(X0) | ! [X1] : (~ p103(X1) | p4(X1) | ~ r1(X0, X1))) & (p4(X0) | ! [X1] : (~ p103(X1) | ~ p4(X1) | ~ r1(X0, X1))))) & (~ p104(X0) | ((~ p5(X0) | ! [X1] : (~ p104(X1) | p5(X1) | ~ r1(X0, X1))) & (p5(X0) | ! [X1] : (~ p104(X1) | ~ p5(X1) | ~ r1(X0, X1))))) & (~ p105(X0) | ((~ p6(X0) | ! [X1] : (~ p105(X1) | p6(X1) | ~ r1(X0, X1))) & (p6(X0) | ! [X1] : (~ p105(X1) | ~ p6(X1) | ~ r1(X0, X1))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & p2(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & ~ p2(X1)) | ~ r1(X0, X1)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & p3(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & ~ p3(X1)) | ~ r1(X0, X1)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & p4(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & ~ p4(X1)) | ~ r1(X0, X1)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & p5(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & ~ p5(X1)) | ~ r1(X0, X1)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & p6(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & ~ p6(X1)) | ~ r1(X0, X1)))) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X1] : (~ p100(X1) | ~ p1(X1) | ~ r1(X0, X1))))) & (~ p101(X0) | ((~ p2(X0) | ! [X1] : (~ p101(X1) | p2(X1) | ~ r1(X0, X1))) & (p2(X0) | ! [X1] : (~ p101(X1) | ~ p2(X1) | ~ r1(X0, X1))))) & (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X1] : (~ p102(X1) | ~ p3(X1) | ~ r1(X0, X1))))) & (~ p103(X0) | ((~ p4(X0) | ! [X1] : (~ p103(X1) | p4(X1) | ~ r1(X0, X1))) & (p4(X0) | ! [X1] : (~ p103(X1) | ~ p4(X1) | ~ r1(X0, X1))))) & (~ p104(X0) | ((~ p5(X0) | ! [X1] : (~ p104(X1) | p5(X1) | ~ r1(X0, X1))) & (p5(X0) | ! [X1] : (~ p104(X1) | ~ p5(X1) | ~ r1(X0, X1))))) & (~ p105(X0) | ((~ p6(X0) | ! [X1] : (~ p105(X1) | p6(X1) | ~ r1(X0, X1))) & (p6(X0) | ! [X1] : (~ p105(X1) | ~ p6(X1) | ~ r1(X0, X1))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & p2(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & ~ p2(X1)) | ~ r1(X0, X1)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & p3(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & ~ p3(X1)) | ~ r1(X0, X1)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & p4(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & ~ p4(X1)) | ~ r1(X0, X1)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & p5(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & ~ p5(X1)) | ~ r1(X0, X1)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & p6(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & ~ p6(X1)) | ~ r1(X0, X1))))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X1] : (~ p100(X1) | ~ p1(X1) | ~ r1(X0, X1))))) & (~ p101(X0) | ((~ p2(X0) | ! [X1] : (~ p101(X1) | p2(X1) | ~ r1(X0, X1))) & (p2(X0) | ! [X1] : (~ p101(X1) | ~ p2(X1) | ~ r1(X0, X1))))) & (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X1] : (~ p102(X1) | ~ p3(X1) | ~ r1(X0, X1))))) & (~ p103(X0) | ((~ p4(X0) | ! [X1] : (~ p103(X1) | p4(X1) | ~ r1(X0, X1))) & (p4(X0) | ! [X1] : (~ p103(X1) | ~ p4(X1) | ~ r1(X0, X1))))) & (~ p104(X0) | ((~ p5(X0) | ! [X1] : (~ p104(X1) | p5(X1) | ~ r1(X0, X1))) & (p5(X0) | ! [X1] : (~ p104(X1) | ~ p5(X1) | ~ r1(X0, X1))))) & (~ p105(X0) | ((~ p6(X0) | ! [X1] : (~ p105(X1) | p6(X1) | ~ r1(X0, X1))) & (p6(X0) | ! [X1] : (~ p105(X1) | ~ p6(X1) | ~ r1(X0, X1))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & p2(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & ~ p2(X1)) | ~ r1(X0, X1)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & p3(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & ~ p3(X1)) | ~ r1(X0, X1)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & p4(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & ~ p4(X1)) | ~ r1(X0, X1)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & p5(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & ~ p5(X1)) | ~ r1(X0, X1)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & p6(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & ~ p6(X1)) | ~ r1(X0, X1))))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & (~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X1] : (~ p100(X1) | ~ p1(X1) | ~ r1(X0, X1))))) & (~ p101(X0) | ((~ p2(X0) | ! [X1] : (~ p101(X1) | p2(X1) | ~ r1(X0, X1))) & (p2(X0) | ! [X1] : (~ p101(X1) | ~ p2(X1) | ~ r1(X0, X1))))) & (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X1] : (~ p102(X1) | ~ p3(X1) | ~ r1(X0, X1))))) & (~ p103(X0) | ((~ p4(X0) | ! [X1] : (~ p103(X1) | p4(X1) | ~ r1(X0, X1))) & (p4(X0) | ! [X1] : (~ p103(X1) | ~ p4(X1) | ~ r1(X0, X1))))) & (~ p104(X0) | ((~ p5(X0) | ! [X1] : (~ p104(X1) | p5(X1) | ~ r1(X0, X1))) & (p5(X0) | ! [X1] : (~ p104(X1) | ~ p5(X1) | ~ r1(X0, X1))))) & (~ p105(X0) | ((~ p6(X0) | ! [X1] : (~ p105(X1) | p6(X1) | ~ r1(X0, X1))) & (p6(X0) | ! [X1] : (~ p105(X1) | ~ p6(X1) | ~ r1(X0, X1))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & p2(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & ~ p2(X1)) | ~ r1(X0, X1)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & p3(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & ~ p3(X1)) | ~ r1(X0, X1)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & p4(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & ~ p4(X1)) | ~ r1(X0, X1)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & p5(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & ~ p5(X1)) | ~ r1(X0, X1)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & p6(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & ~ p6(X1)) | ~ r1(X0, X1)))) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X1] : (~ p100(X1) | ~ p1(X1) | ~ r1(X0, X1))))) & (~ p101(X0) | ((~ p2(X0) | ! [X1] : (~ p101(X1) | p2(X1) | ~ r1(X0, X1))) & (p2(X0) | ! [X1] : (~ p101(X1) | ~ p2(X1) | ~ r1(X0, X1))))) & (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X1] : (~ p102(X1) | ~ p3(X1) | ~ r1(X0, X1))))) & (~ p103(X0) | ((~ p4(X0) | ! [X1] : (~ p103(X1) | p4(X1) | ~ r1(X0, X1))) & (p4(X0) | ! [X1] : (~ p103(X1) | ~ p4(X1) | ~ r1(X0, X1))))) & (~ p104(X0) | ((~ p5(X0) | ! [X1] : (~ p104(X1) | p5(X1) | ~ r1(X0, X1))) & (p5(X0) | ! [X1] : (~ p104(X1) | ~ p5(X1) | ~ r1(X0, X1))))) & (~ p105(X0) | ((~ p6(X0) | ! [X1] : (~ p105(X1) | p6(X1) | ~ r1(X0, X1))) & (p6(X0) | ! [X1] : (~ p105(X1) | ~ p6(X1) | ~ r1(X0, X1))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & p2(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & ~ p2(X1)) | ~ r1(X0, X1)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & p3(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & ~ p3(X1)) | ~ r1(X0, X1)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & p4(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & ~ p4(X1)) | ~ r1(X0, X1)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & p5(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & ~ p5(X1)) | ~ r1(X0, X1)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & p6(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & ~ p6(X1)) | ~ r1(X0, X1))))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X1] : (~ p100(X1) | ~ p1(X1) | ~ r1(X0, X1))))) & (~ p101(X0) | ((~ p2(X0) | ! [X1] : (~ p101(X1) | p2(X1) | ~ r1(X0, X1))) & (p2(X0) | ! [X1] : (~ p101(X1) | ~ p2(X1) | ~ r1(X0, X1))))) & (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X1] : (~ p102(X1) | ~ p3(X1) | ~ r1(X0, X1))))) & (~ p103(X0) | ((~ p4(X0) | ! [X1] : (~ p103(X1) | p4(X1) | ~ r1(X0, X1))) & (p4(X0) | ! [X1] : (~ p103(X1) | ~ p4(X1) | ~ r1(X0, X1))))) & (~ p104(X0) | ((~ p5(X0) | ! [X1] : (~ p104(X1) | p5(X1) | ~ r1(X0, X1))) & (p5(X0) | ! [X1] : (~ p104(X1) | ~ p5(X1) | ~ r1(X0, X1))))) & (~ p105(X0) | ((~ p6(X0) | ! [X1] : (~ p105(X1) | p6(X1) | ~ r1(X0, X1))) & (p6(X0) | ! [X1] : (~ p105(X1) | ~ p6(X1) | ~ r1(X0, X1))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & p2(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & ~ p2(X1)) | ~ r1(X0, X1)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & p3(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & ~ p3(X1)) | ~ r1(X0, X1)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & p4(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & ~ p4(X1)) | ~ r1(X0, X1)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & p5(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & ~ p5(X1)) | ~ r1(X0, X1)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & p6(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & ~ p6(X1)) | ~ r1(X0, X1))))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL636+1.005.p', main)).
fof(f694, plain, (p101(sK110) | r1(sK110, sK111)), inference(subsumption_resolution, [], [f647, f669])).
fof(f669, plain, p100(sK110), inference(cnf_transformation, [], [f276])).
fof(f647, plain, (~ p100(sK110) | p101(sK110) | r1(sK110, sK111)), inference(cnf_transformation, [], [f276])).
fof(f612, plain, ! [X23] : (~ r1(sK110, X23) | sP59(X23)), inference(cnf_transformation, [], [f276])).
fof(f21279, plain, (~ sP50(sK111) | ~ spl121_188), inference(subsumption_resolution, [], [f21278, f689])).
fof(f689, plain, p101(sK111), inference(subsumption_resolution, [], [f688, f668])).
fof(f688, plain, (p101(sK110) | p101(sK111)), inference(subsumption_resolution, [], [f650, f669])).
fof(f650, plain, (~ p100(sK110) | p101(sK110) | p101(sK111)), inference(cnf_transformation, [], [f276])).
fof(f21278, plain, (~ p101(sK111) | ~ sP50(sK111) | ~ spl121_188), inference(subsumption_resolution, [], [f21277, f691])).
fof(f691, plain, ~ p102(sK111), inference(subsumption_resolution, [], [f690, f668])).
fof(f690, plain, (p101(sK110) | ~ p102(sK111)), inference(subsumption_resolution, [], [f649, f669])).
fof(f649, plain, (~ p100(sK110) | p101(sK110) | ~ p102(sK111)), inference(cnf_transformation, [], [f276])).
fof(f21277, plain, (p102(sK111) | ~ p101(sK111) | ~ sP50(sK111) | ~ spl121_188), inference(resolution, [], [f6805, f320])).
fof(f320, plain, ! [X0] : (~ p3(sK65(X0)) | p102(X0) | ~ p101(X0) | ~ sP50(X0)), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ! [X0] : (~ p101(X0) | p102(X0) | ((p102(sK64(X0)) & ~ p103(sK64(X0)) & p3(sK64(X0)) & r1(X0, sK64(X0))) & (p102(sK65(X0)) & ~ p103(sK65(X0)) & ~ p3(sK65(X0)) & r1(X0, sK65(X0)))) | ~ sP50(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK64, sK65])], [f94, f96, f95])).
fof(f95, plain, ! [X0] : (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) => (p102(sK64(X0)) & ~ p103(sK64(X0)) & p3(sK64(X0)) & r1(X0, sK64(X0)))), introduced(choice_axiom, [])).
fof(f96, plain, ! [X0] : (? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2)) => (p102(sK65(X0)) & ~ p103(sK65(X0)) & ~ p3(sK65(X0)) & r1(X0, sK65(X0)))), introduced(choice_axiom, [])).
fof(f94, plain, ! [X0] : (~ p101(X0) | p102(X0) | (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) & ? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2))) | ~ sP50(X0)), inference(rectify, [], [f93])).
fof(f93, plain, ! [X23] : (~ p101(X23) | p102(X23) | (? [X38] : (p102(X38) & ~ p103(X38) & p3(X38) & r1(X23, X38)) & ? [X39] : (p102(X39) & ~ p103(X39) & ~ p3(X39) & r1(X23, X39))) | ~ sP50(X23)), inference(nnf_transformation, [], [f58])).
fof(f6805, plain, (p3(sK65(sK111)) | ~ spl121_188), inference(avatar_component_clause, [], [f6804])).
fof(f6804, plain, (spl121_188 <=> p3(sK65(sK111))), introduced(avatar_definition, [new_symbols(naming, [spl121_188])])).
fof(f21216, plain, (~ spl121_633 | spl121_247 | spl121_148 | ~ spl121_149 | ~ spl121_150 | ~ spl121_167 | ~ spl121_459), inference(avatar_split_clause, [], [f21215, f9733, f6580, f6473, f6468, f6463, f7383, f11398])).
fof(f11398, plain, (spl121_633 <=> p3(sK89(sK77(sK65(sK111))))), introduced(avatar_definition, [new_symbols(naming, [spl121_633])])).
fof(f7383, plain, (spl121_247 <=> p3(sK77(sK65(sK111)))), introduced(avatar_definition, [new_symbols(naming, [spl121_247])])).
fof(f6463, plain, (spl121_148 <=> p104(sK77(sK65(sK111)))), introduced(avatar_definition, [new_symbols(naming, [spl121_148])])).
fof(f6468, plain, (spl121_149 <=> p103(sK77(sK65(sK111)))), introduced(avatar_definition, [new_symbols(naming, [spl121_149])])).
fof(f6473, plain, (spl121_150 <=> p102(sK77(sK65(sK111)))), introduced(avatar_definition, [new_symbols(naming, [spl121_150])])).
fof(f6580, plain, (spl121_167 <=> sP35(sK77(sK65(sK111)))), introduced(avatar_definition, [new_symbols(naming, [spl121_167])])).
fof(f9733, plain, (spl121_459 <=> p102(sK89(sK77(sK65(sK111))))), introduced(avatar_definition, [new_symbols(naming, [spl121_459])])).
fof(f21215, plain, (p3(sK77(sK65(sK111))) | ~ p3(sK89(sK77(sK65(sK111)))) | (spl121_148 | ~ spl121_149 | ~ spl121_150 | ~ spl121_167 | ~ spl121_459)), inference(subsumption_resolution, [], [f21214, f6759])).
fof(f6759, plain, (sP24(sK77(sK65(sK111))) | ~ spl121_167), inference(resolution, [], [f6582, f410])).
fof(f410, plain, ! [X0] : (~ sP35(X0) | sP24(X0)), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & sP34(X0) & sP33(X0) & sP32(X0) & sP31(X0) & sP30(X0) & sP29(X0) & sP27(X0) & sP26(X0) & sP25(X0) & sP24(X0) & sP28(X0)) | ~ sP35(X0)), inference(rectify, [], [f147])).
fof(f147, plain, ! [X72] : (((~ p101(X72) | p100(X72)) & (~ p102(X72) | p101(X72)) & (~ p103(X72) | p102(X72)) & (~ p104(X72) | p103(X72)) & (~ p105(X72) | p104(X72)) & sP34(X72) & sP33(X72) & sP32(X72) & sP31(X72) & sP30(X72) & sP29(X72) & sP27(X72) & sP26(X72) & sP25(X72) & sP24(X72) & sP28(X72)) | ~ sP35(X72)), inference(nnf_transformation, [], [f43])).
fof(f6582, plain, (sP35(sK77(sK65(sK111))) | ~ spl121_167), inference(avatar_component_clause, [], [f6580])).
fof(f21214, plain, (p3(sK77(sK65(sK111))) | ~ p3(sK89(sK77(sK65(sK111)))) | ~ sP24(sK77(sK65(sK111))) | (spl121_148 | ~ spl121_149 | ~ spl121_150 | ~ spl121_167 | ~ spl121_459)), inference(subsumption_resolution, [], [f21213, f6470])).
fof(f6470, plain, (p103(sK77(sK65(sK111))) | ~ spl121_149), inference(avatar_component_clause, [], [f6468])).
fof(f21213, plain, (p3(sK77(sK65(sK111))) | ~ p3(sK89(sK77(sK65(sK111)))) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111))) | (spl121_148 | ~ spl121_150 | ~ spl121_167 | ~ spl121_459)), inference(subsumption_resolution, [], [f21212, f6464])).
fof(f6464, plain, (~ p104(sK77(sK65(sK111))) | spl121_148), inference(avatar_component_clause, [], [f6463])).
fof(f21212, plain, (p3(sK77(sK65(sK111))) | ~ p3(sK89(sK77(sK65(sK111)))) | p104(sK77(sK65(sK111))) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111))) | (~ spl121_150 | ~ spl121_167 | ~ spl121_459)), inference(subsumption_resolution, [], [f21211, f6766])).
fof(f6766, plain, (sP32(sK77(sK65(sK111))) | ~ spl121_167), inference(resolution, [], [f6582, f417])).
fof(f417, plain, ! [X0] : (~ sP35(X0) | sP32(X0)), inference(cnf_transformation, [], [f148])).
fof(f21211, plain, (p3(sK77(sK65(sK111))) | ~ p3(sK89(sK77(sK65(sK111)))) | ~ sP32(sK77(sK65(sK111))) | p104(sK77(sK65(sK111))) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111))) | (~ spl121_150 | ~ spl121_459)), inference(subsumption_resolution, [], [f18412, f6475])).
fof(f6475, plain, (p102(sK77(sK65(sK111))) | ~ spl121_150), inference(avatar_component_clause, [], [f6473])).
fof(f18412, plain, (p3(sK77(sK65(sK111))) | ~ p3(sK89(sK77(sK65(sK111)))) | ~ p102(sK77(sK65(sK111))) | ~ sP32(sK77(sK65(sK111))) | p104(sK77(sK65(sK111))) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111))) | ~ spl121_459), inference(resolution, [], [f3242, f9735])).
fof(f9735, plain, (p102(sK89(sK77(sK65(sK111)))) | ~ spl121_459), inference(avatar_component_clause, [], [f9733])).
fof(f3242, plain, ! [X29] : (~ p102(sK89(X29)) | p3(X29) | ~ p3(sK89(X29)) | ~ p102(X29) | ~ sP32(X29) | p104(X29) | ~ p103(X29) | ~ sP24(X29)), inference(resolution, [], [f429, f467])).
fof(f467, plain, ! [X0] : (r1(X0, sK89(X0)) | p104(X0) | ~ p103(X0) | ~ sP24(X0)), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ! [X0] : (~ p103(X0) | p104(X0) | ((p104(sK88(X0)) & ~ p105(sK88(X0)) & p5(sK88(X0)) & r1(X0, sK88(X0))) & (p104(sK89(X0)) & ~ p105(sK89(X0)) & ~ p5(sK89(X0)) & r1(X0, sK89(X0)))) | ~ sP24(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK88, sK89])], [f182, f184, f183])).
fof(f183, plain, ! [X0] : (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) => (p104(sK88(X0)) & ~ p105(sK88(X0)) & p5(sK88(X0)) & r1(X0, sK88(X0)))), introduced(choice_axiom, [])).
fof(f184, plain, ! [X0] : (? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2)) => (p104(sK89(X0)) & ~ p105(sK89(X0)) & ~ p5(sK89(X0)) & r1(X0, sK89(X0)))), introduced(choice_axiom, [])).
fof(f182, plain, ! [X0] : (~ p103(X0) | p104(X0) | (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) & ? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2))) | ~ sP24(X0)), inference(rectify, [], [f181])).
fof(f181, plain, ! [X72] : (~ p103(X72) | p104(X72) | (? [X91] : (p104(X91) & ~ p105(X91) & p5(X91) & r1(X72, X91)) & ? [X92] : (p104(X92) & ~ p105(X92) & ~ p5(X92) & r1(X72, X92))) | ~ sP24(X72)), inference(nnf_transformation, [], [f32])).
fof(f429, plain, ! [X2, X0] : (~ r1(X0, X2) | p3(X0) | ~ p102(X2) | ~ p3(X2) | ~ p102(X0) | ~ sP32(X0)), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ! [X0] : (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X2] : (~ p102(X2) | ~ p3(X2) | ~ r1(X0, X2)))) | ~ sP32(X0)), inference(rectify, [], [f153])).
fof(f153, plain, ! [X72] : (~ p102(X72) | ((~ p3(X72) | ! [X77] : (~ p102(X77) | p3(X77) | ~ r1(X72, X77))) & (p3(X72) | ! [X78] : (~ p102(X78) | ~ p3(X78) | ~ r1(X72, X78)))) | ~ sP32(X72)), inference(nnf_transformation, [], [f40])).
fof(f21157, plain, (spl121_633 | spl121_456 | ~ spl121_457 | ~ spl121_459 | ~ spl121_468 | ~ spl121_1225 | ~ spl121_1237), inference(avatar_split_clause, [], [f21156, f19764, f19673, f9796, f9733, f9723, f9719, f11398])).
fof(f9719, plain, (spl121_456 <=> p105(sK89(sK77(sK65(sK111))))), introduced(avatar_definition, [new_symbols(naming, [spl121_456])])).
fof(f9723, plain, (spl121_457 <=> p104(sK89(sK77(sK65(sK111))))), introduced(avatar_definition, [new_symbols(naming, [spl121_457])])).
fof(f9796, plain, (spl121_468 <=> sP23(sK89(sK77(sK65(sK111))))), introduced(avatar_definition, [new_symbols(naming, [spl121_468])])).
fof(f19673, plain, (spl121_1225 <=> p102(sK90(sK89(sK77(sK65(sK111)))))), introduced(avatar_definition, [new_symbols(naming, [spl121_1225])])).
fof(f19764, plain, (spl121_1237 <=> p3(sK90(sK89(sK77(sK65(sK111)))))), introduced(avatar_definition, [new_symbols(naming, [spl121_1237])])).
fof(f21156, plain, (p3(sK89(sK77(sK65(sK111)))) | (spl121_456 | ~ spl121_457 | ~ spl121_459 | ~ spl121_468 | ~ spl121_1225 | ~ spl121_1237)), inference(subsumption_resolution, [], [f21155, f9893])).
fof(f9893, plain, (sP16(sK89(sK77(sK65(sK111)))) | ~ spl121_468), inference(resolution, [], [f9798, f475])).
fof(f475, plain, ! [X0] : (~ sP23(X0) | sP16(X0)), inference(cnf_transformation, [], [f187])).
fof(f187, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & sP22(X0) & sP21(X0) & sP20(X0) & sP19(X0) & sP18(X0) & sP17(X0) & sP15(X0) & sP14(X0) & sP13(X0) & sP12(X0) & sP16(X0)) | ~ sP23(X0)), inference(rectify, [], [f186])).
fof(f186, plain, ! [X98] : (((~ p101(X98) | p100(X98)) & (~ p102(X98) | p101(X98)) & (~ p103(X98) | p102(X98)) & (~ p104(X98) | p103(X98)) & (~ p105(X98) | p104(X98)) & sP22(X98) & sP21(X98) & sP20(X98) & sP19(X98) & sP18(X98) & sP17(X98) & sP15(X98) & sP14(X98) & sP13(X98) & sP12(X98) & sP16(X98)) | ~ sP23(X98)), inference(nnf_transformation, [], [f31])).
fof(f9798, plain, (sP23(sK89(sK77(sK65(sK111)))) | ~ spl121_468), inference(avatar_component_clause, [], [f9796])).
fof(f21155, plain, (p3(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (spl121_456 | ~ spl121_457 | ~ spl121_459 | ~ spl121_468 | ~ spl121_1225 | ~ spl121_1237)), inference(subsumption_resolution, [], [f21154, f9725])).
fof(f9725, plain, (p104(sK89(sK77(sK65(sK111)))) | ~ spl121_457), inference(avatar_component_clause, [], [f9723])).
fof(f21154, plain, (p3(sK89(sK77(sK65(sK111)))) | ~ p104(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (spl121_456 | ~ spl121_459 | ~ spl121_468 | ~ spl121_1225 | ~ spl121_1237)), inference(subsumption_resolution, [], [f21153, f9721])).
fof(f9721, plain, (~ p105(sK89(sK77(sK65(sK111)))) | spl121_456), inference(avatar_component_clause, [], [f9719])).
fof(f21153, plain, (p3(sK89(sK77(sK65(sK111)))) | p105(sK89(sK77(sK65(sK111)))) | ~ p104(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (~ spl121_459 | ~ spl121_468 | ~ spl121_1225 | ~ spl121_1237)), inference(subsumption_resolution, [], [f21152, f9901])).
fof(f9901, plain, (sP20(sK89(sK77(sK65(sK111)))) | ~ spl121_468), inference(resolution, [], [f9798, f483])).
fof(f483, plain, ! [X0] : (~ sP23(X0) | sP20(X0)), inference(cnf_transformation, [], [f187])).
fof(f21152, plain, (p3(sK89(sK77(sK65(sK111)))) | ~ sP20(sK89(sK77(sK65(sK111)))) | p105(sK89(sK77(sK65(sK111)))) | ~ p104(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (~ spl121_459 | ~ spl121_1225 | ~ spl121_1237)), inference(subsumption_resolution, [], [f21151, f9735])).
fof(f21151, plain, (p3(sK89(sK77(sK65(sK111)))) | ~ p102(sK89(sK77(sK65(sK111)))) | ~ sP20(sK89(sK77(sK65(sK111)))) | p105(sK89(sK77(sK65(sK111)))) | ~ p104(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (~ spl121_1225 | ~ spl121_1237)), inference(subsumption_resolution, [], [f20858, f19766])).
fof(f19766, plain, (p3(sK90(sK89(sK77(sK65(sK111))))) | ~ spl121_1237), inference(avatar_component_clause, [], [f19764])).
fof(f20858, plain, (p3(sK89(sK77(sK65(sK111)))) | ~ p3(sK90(sK89(sK77(sK65(sK111))))) | ~ p102(sK89(sK77(sK65(sK111)))) | ~ sP20(sK89(sK77(sK65(sK111)))) | p105(sK89(sK77(sK65(sK111)))) | ~ p104(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | ~ spl121_1225), inference(resolution, [], [f3972, f19675])).
fof(f19675, plain, (p102(sK90(sK89(sK77(sK65(sK111))))) | ~ spl121_1225), inference(avatar_component_clause, [], [f19673])).
fof(f3972, plain, ! [X30] : (~ p102(sK90(X30)) | p3(X30) | ~ p3(sK90(X30)) | ~ p102(X30) | ~ sP20(X30) | p105(X30) | ~ p104(X30) | ~ sP16(X30)), inference(resolution, [], [f495, f506])).
fof(f506, plain, ! [X0] : (r1(X0, sK90(X0)) | p105(X0) | ~ p104(X0) | ~ sP16(X0)), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ! [X0] : (~ p104(X0) | p105(X0) | ((p105(sK90(X0)) & p6(sK90(X0)) & r1(X0, sK90(X0))) & (p105(sK91(X0)) & ~ p6(sK91(X0)) & r1(X0, sK91(X0)))) | ~ sP16(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK90, sK91])], [f201, f203, f202])).
fof(f202, plain, ! [X0] : (? [X1] : (p105(X1) & p6(X1) & r1(X0, X1)) => (p105(sK90(X0)) & p6(sK90(X0)) & r1(X0, sK90(X0)))), introduced(choice_axiom, [])).
fof(f203, plain, ! [X0] : (? [X2] : (p105(X2) & ~ p6(X2) & r1(X0, X2)) => (p105(sK91(X0)) & ~ p6(sK91(X0)) & r1(X0, sK91(X0)))), introduced(choice_axiom, [])).
fof(f201, plain, ! [X0] : (~ p104(X0) | p105(X0) | (? [X1] : (p105(X1) & p6(X1) & r1(X0, X1)) & ? [X2] : (p105(X2) & ~ p6(X2) & r1(X0, X2))) | ~ sP16(X0)), inference(rectify, [], [f200])).
fof(f200, plain, ! [X98] : (~ p104(X98) | p105(X98) | (? [X119] : (p105(X119) & p6(X119) & r1(X98, X119)) & ? [X120] : (p105(X120) & ~ p6(X120) & r1(X98, X120))) | ~ sP16(X98)), inference(nnf_transformation, [], [f24])).
fof(f495, plain, ! [X2, X0] : (~ r1(X0, X2) | p3(X0) | ~ p102(X2) | ~ p3(X2) | ~ p102(X0) | ~ sP20(X0)), inference(cnf_transformation, [], [f193])).
fof(f193, plain, ! [X0] : (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X2] : (~ p102(X2) | ~ p3(X2) | ~ r1(X0, X2)))) | ~ sP20(X0)), inference(rectify, [], [f192])).
fof(f192, plain, ! [X98] : (~ p102(X98) | ((~ p3(X98) | ! [X103] : (~ p102(X103) | p3(X103) | ~ r1(X98, X103))) & (p3(X98) | ! [X104] : (~ p102(X104) | ~ p3(X104) | ~ r1(X98, X104)))) | ~ sP20(X98)), inference(nnf_transformation, [], [f28])).
fof(f19772, plain, (spl121_148 | ~ spl121_149 | ~ spl121_167 | ~ spl121_456), inference(avatar_contradiction_clause, [], [f19771])).
fof(f19771, plain, ($false | (spl121_148 | ~ spl121_149 | ~ spl121_167 | ~ spl121_456)), inference(subsumption_resolution, [], [f19770, f6759])).
fof(f19770, plain, (~ sP24(sK77(sK65(sK111))) | (spl121_148 | ~ spl121_149 | ~ spl121_456)), inference(subsumption_resolution, [], [f19769, f6470])).
fof(f19769, plain, (~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111))) | (spl121_148 | ~ spl121_456)), inference(subsumption_resolution, [], [f19768, f6464])).
fof(f19768, plain, (p104(sK77(sK65(sK111))) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111))) | ~ spl121_456), inference(resolution, [], [f9720, f469])).
fof(f469, plain, ! [X0] : (~ p105(sK89(X0)) | p104(X0) | ~ p103(X0) | ~ sP24(X0)), inference(cnf_transformation, [], [f185])).
fof(f9720, plain, (p105(sK89(sK77(sK65(sK111)))) | ~ spl121_456), inference(avatar_component_clause, [], [f9719])).
fof(f19767, plain, (spl121_456 | spl121_1237 | spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | ~ spl121_457 | ~ spl121_468), inference(avatar_split_clause, [], [f19762, f9796, f9723, f6580, f6468, f6463, f868, f863, f19764, f9719])).
fof(f863, plain, (spl121_17 <=> p103(sK65(sK111))), introduced(avatar_definition, [new_symbols(naming, [spl121_17])])).
fof(f868, plain, (spl121_18 <=> p102(sK65(sK111))), introduced(avatar_definition, [new_symbols(naming, [spl121_18])])).
fof(f19762, plain, (p3(sK90(sK89(sK77(sK65(sK111))))) | p105(sK89(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | ~ spl121_457 | ~ spl121_468)), inference(subsumption_resolution, [], [f19761, f9893])).
fof(f19761, plain, (p3(sK90(sK89(sK77(sK65(sK111))))) | p105(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | ~ spl121_457)), inference(subsumption_resolution, [], [f13849, f9725])).
fof(f13849, plain, (p3(sK90(sK89(sK77(sK65(sK111))))) | p105(sK89(sK77(sK65(sK111)))) | ~ p104(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167)), inference(resolution, [], [f13686, f506])).
fof(f13686, plain, (! [X29] : (~ r1(sK89(sK77(sK65(sK111))), X29) | p3(X29)) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167)), inference(subsumption_resolution, [], [f13685, f6759])).
fof(f13685, plain, (! [X29] : (~ r1(sK89(sK77(sK65(sK111))), X29) | p3(X29) | ~ sP24(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149)), inference(subsumption_resolution, [], [f13684, f6470])).
fof(f13684, plain, (! [X29] : (~ r1(sK89(sK77(sK65(sK111))), X29) | p3(X29) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18 | spl121_148)), inference(subsumption_resolution, [], [f13660, f6464])).
fof(f13660, plain, (! [X29] : (~ r1(sK89(sK77(sK65(sK111))), X29) | p3(X29) | p104(sK77(sK65(sK111))) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18)), inference(resolution, [], [f13318, f467])).
fof(f13318, plain, (! [X35, X34] : (~ r1(sK77(sK65(sK111)), X34) | ~ r1(X34, X35) | p3(X35)) | (spl121_17 | ~ spl121_18)), inference(subsumption_resolution, [], [f13317, f839])).
fof(f839, plain, sP37(sK65(sK111)), inference(resolution, [], [f833, f345])).
fof(f345, plain, ! [X0] : (~ sP47(X0) | sP37(X0)), inference(cnf_transformation, [], [f109])).
fof(f109, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & sP46(X0) & sP45(X0) & sP44(X0) & sP43(X0) & sP42(X0) & sP41(X0) & sP39(X0) & sP38(X0) & sP37(X0) & sP36(X0) & sP40(X0)) | ~ sP47(X0)), inference(rectify, [], [f108])).
fof(f108, plain, ! [X47] : (((~ p101(X47) | p100(X47)) & (~ p102(X47) | p101(X47)) & (~ p103(X47) | p102(X47)) & (~ p104(X47) | p103(X47)) & (~ p105(X47) | p104(X47)) & sP46(X47) & sP45(X47) & sP44(X47) & sP43(X47) & sP42(X47) & sP41(X47) & sP39(X47) & sP38(X47) & sP37(X47) & sP36(X47) & sP40(X47)) | ~ sP47(X47)), inference(nnf_transformation, [], [f55])).
fof(f833, plain, sP47(sK65(sK111)), inference(subsumption_resolution, [], [f832, f709])).
fof(f832, plain, (~ sP50(sK111) | sP47(sK65(sK111))), inference(subsumption_resolution, [], [f831, f689])).
fof(f831, plain, (~ p101(sK111) | ~ sP50(sK111) | sP47(sK65(sK111))), inference(subsumption_resolution, [], [f829, f691])).
fof(f829, plain, (p102(sK111) | ~ p101(sK111) | ~ sP50(sK111) | sP47(sK65(sK111))), inference(resolution, [], [f319, f776])).
fof(f776, plain, ! [X0] : (~ r1(sK111, X0) | sP47(X0)), inference(resolution, [], [f611, f695])).
fof(f611, plain, ! [X24, X25] : (~ r1(sK110, X24) | ~ r1(X24, X25) | sP47(X25)), inference(cnf_transformation, [], [f276])).
fof(f319, plain, ! [X0] : (r1(X0, sK65(X0)) | p102(X0) | ~ p101(X0) | ~ sP50(X0)), inference(cnf_transformation, [], [f97])).
fof(f13317, plain, (! [X35, X34] : (~ r1(sK77(sK65(sK111)), X34) | ~ r1(X34, X35) | p3(X35) | ~ sP37(sK65(sK111))) | (spl121_17 | ~ spl121_18)), inference(subsumption_resolution, [], [f13316, f870])).
fof(f870, plain, (p102(sK65(sK111)) | ~ spl121_18), inference(avatar_component_clause, [], [f868])).
fof(f13316, plain, (! [X35, X34] : (~ r1(sK77(sK65(sK111)), X34) | ~ r1(X34, X35) | p3(X35) | ~ p102(sK65(sK111)) | ~ sP37(sK65(sK111))) | spl121_17), inference(subsumption_resolution, [], [f13280, f864])).
fof(f864, plain, (~ p103(sK65(sK111)) | spl121_17), inference(avatar_component_clause, [], [f863])).
fof(f13280, plain, ! [X35, X34] : (~ r1(sK77(sK65(sK111)), X34) | ~ r1(X34, X35) | p3(X35) | p103(sK65(sK111)) | ~ p102(sK65(sK111)) | ~ sP37(sK65(sK111))), inference(resolution, [], [f6044, f393])).
fof(f393, plain, ! [X0] : (r1(X0, sK77(X0)) | p103(X0) | ~ p102(X0) | ~ sP37(X0)), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ! [X0] : (~ p102(X0) | p103(X0) | ((p103(sK76(X0)) & ~ p104(sK76(X0)) & p4(sK76(X0)) & r1(X0, sK76(X0))) & (p103(sK77(X0)) & ~ p104(sK77(X0)) & ~ p4(sK77(X0)) & r1(X0, sK77(X0)))) | ~ sP37(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK76, sK77])], [f138, f140, f139])).
fof(f139, plain, ! [X0] : (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) => (p103(sK76(X0)) & ~ p104(sK76(X0)) & p4(sK76(X0)) & r1(X0, sK76(X0)))), introduced(choice_axiom, [])).
fof(f140, plain, ! [X0] : (? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2)) => (p103(sK77(X0)) & ~ p104(sK77(X0)) & ~ p4(sK77(X0)) & r1(X0, sK77(X0)))), introduced(choice_axiom, [])).
fof(f138, plain, ! [X0] : (~ p102(X0) | p103(X0) | (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) & ? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2))) | ~ sP37(X0)), inference(rectify, [], [f137])).
fof(f137, plain, ! [X47] : (~ p102(X47) | p103(X47) | (? [X64] : (p103(X64) & ~ p104(X64) & p4(X64) & r1(X47, X64)) & ? [X65] : (p103(X65) & ~ p104(X65) & ~ p4(X65) & r1(X47, X65))) | ~ sP37(X47)), inference(nnf_transformation, [], [f45])).
fof(f6044, plain, ! [X17, X15, X16] : (~ r1(sK65(sK111), X15) | ~ r1(X15, X16) | ~ r1(X16, X17) | p3(X17)), inference(subsumption_resolution, [], [f6043, f709])).
fof(f6043, plain, ! [X17, X15, X16] : (~ r1(X15, X16) | ~ r1(sK65(sK111), X15) | ~ r1(X16, X17) | p3(X17) | ~ sP50(sK111)), inference(subsumption_resolution, [], [f6042, f689])).
fof(f6042, plain, ! [X17, X15, X16] : (~ r1(X15, X16) | ~ r1(sK65(sK111), X15) | ~ r1(X16, X17) | p3(X17) | ~ p101(sK111) | ~ sP50(sK111)), inference(subsumption_resolution, [], [f5994, f691])).
fof(f5994, plain, ! [X17, X15, X16] : (~ r1(X15, X16) | ~ r1(sK65(sK111), X15) | ~ r1(X16, X17) | p3(X17) | p102(sK111) | ~ p101(sK111) | ~ sP50(sK111)), inference(resolution, [], [f5466, f319])).
fof(f5466, plain, ! [X2, X0, X3, X1] : (~ r1(sK111, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | p3(X1)), inference(resolution, [], [f607, f695])).
fof(f607, plain, ! [X39, X41, X38, X42, X40] : (~ r1(sK110, X38) | ~ r1(X41, X42) | ~ r1(X40, X41) | ~ r1(X39, X40) | ~ r1(X38, X39) | p3(X42)), inference(cnf_transformation, [], [f276])).
fof(f19739, plain, (spl121_456 | ~ spl121_457 | ~ spl121_468 | spl121_1222), inference(avatar_split_clause, [], [f19738, f19659, f9796, f9723, f9719])).
fof(f19659, plain, (spl121_1222 <=> p105(sK90(sK89(sK77(sK65(sK111)))))), introduced(avatar_definition, [new_symbols(naming, [spl121_1222])])).
fof(f19738, plain, (p105(sK89(sK77(sK65(sK111)))) | (~ spl121_457 | ~ spl121_468 | spl121_1222)), inference(subsumption_resolution, [], [f19737, f9893])).
fof(f19737, plain, (p105(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (~ spl121_457 | spl121_1222)), inference(subsumption_resolution, [], [f19732, f9725])).
fof(f19732, plain, (p105(sK89(sK77(sK65(sK111)))) | ~ p104(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | spl121_1222), inference(resolution, [], [f19661, f508])).
fof(f508, plain, ! [X0] : (p105(sK90(X0)) | p105(X0) | ~ p104(X0) | ~ sP16(X0)), inference(cnf_transformation, [], [f204])).
fof(f19661, plain, (~ p105(sK90(sK89(sK77(sK65(sK111))))) | spl121_1222), inference(avatar_component_clause, [], [f19659])).
fof(f19676, plain, (~ spl121_1224 | spl121_1225 | spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_456 | ~ spl121_457 | ~ spl121_468), inference(avatar_split_clause, [], [f19655, f9796, f9723, f9719, f6580, f6468, f6463, f868, f863, f19673, f19668])).
fof(f19668, plain, (spl121_1224 <=> p103(sK90(sK89(sK77(sK65(sK111)))))), introduced(avatar_definition, [new_symbols(naming, [spl121_1224])])).
fof(f19655, plain, (p102(sK90(sK89(sK77(sK65(sK111))))) | ~ p103(sK90(sK89(sK77(sK65(sK111))))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_456 | ~ spl121_457 | ~ spl121_468)), inference(resolution, [], [f19638, f554])).
fof(f554, plain, ! [X0] : (~ sP11(X0) | p102(X0) | ~ p103(X0)), inference(cnf_transformation, [], [f226])).
fof(f226, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & sP10(X0) & sP9(X0) & sP8(X0) & sP7(X0) & sP6(X0) & sP5(X0) & sP3(X0) & sP2(X0) & sP1(X0) & sP0(X0) & sP4(X0)) | ~ sP11(X0)), inference(rectify, [], [f225])).
fof(f225, plain, ! [X125] : (((~ p101(X125) | p100(X125)) & (~ p102(X125) | p101(X125)) & (~ p103(X125) | p102(X125)) & (~ p104(X125) | p103(X125)) & (~ p105(X125) | p104(X125)) & sP10(X125) & sP9(X125) & sP8(X125) & sP7(X125) & sP6(X125) & sP5(X125) & sP3(X125) & sP2(X125) & sP1(X125) & sP0(X125) & sP4(X125)) | ~ sP11(X125)), inference(nnf_transformation, [], [f19])).
fof(f19638, plain, (sP11(sK90(sK89(sK77(sK65(sK111))))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_456 | ~ spl121_457 | ~ spl121_468)), inference(subsumption_resolution, [], [f19637, f9893])).
fof(f19637, plain, (sP11(sK90(sK89(sK77(sK65(sK111))))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_456 | ~ spl121_457)), inference(subsumption_resolution, [], [f19636, f9725])).
fof(f19636, plain, (sP11(sK90(sK89(sK77(sK65(sK111))))) | ~ p104(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_456)), inference(subsumption_resolution, [], [f19616, f9721])).
fof(f19616, plain, (sP11(sK90(sK89(sK77(sK65(sK111))))) | p105(sK89(sK77(sK65(sK111)))) | ~ p104(sK89(sK77(sK65(sK111)))) | ~ sP16(sK89(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167)), inference(resolution, [], [f19112, f506])).
fof(f19112, plain, (! [X29] : (~ r1(sK89(sK77(sK65(sK111))), X29) | sP11(X29)) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167)), inference(subsumption_resolution, [], [f19111, f6759])).
fof(f19111, plain, (! [X29] : (~ r1(sK89(sK77(sK65(sK111))), X29) | sP11(X29) | ~ sP24(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149)), inference(subsumption_resolution, [], [f19110, f6470])).
fof(f19110, plain, (! [X29] : (~ r1(sK89(sK77(sK65(sK111))), X29) | sP11(X29) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18 | spl121_148)), inference(subsumption_resolution, [], [f19086, f6464])).
fof(f19086, plain, (! [X29] : (~ r1(sK89(sK77(sK65(sK111))), X29) | sP11(X29) | p104(sK77(sK65(sK111))) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111)))) | (spl121_17 | ~ spl121_18)), inference(resolution, [], [f17621, f467])).
fof(f17621, plain, (! [X35, X34] : (~ r1(sK77(sK65(sK111)), X34) | ~ r1(X34, X35) | sP11(X35)) | (spl121_17 | ~ spl121_18)), inference(subsumption_resolution, [], [f17620, f839])).
fof(f17620, plain, (! [X35, X34] : (~ r1(sK77(sK65(sK111)), X34) | ~ r1(X34, X35) | sP11(X35) | ~ sP37(sK65(sK111))) | (spl121_17 | ~ spl121_18)), inference(subsumption_resolution, [], [f17619, f870])).
fof(f17619, plain, (! [X35, X34] : (~ r1(sK77(sK65(sK111)), X34) | ~ r1(X34, X35) | sP11(X35) | ~ p102(sK65(sK111)) | ~ sP37(sK65(sK111))) | spl121_17), inference(subsumption_resolution, [], [f17583, f864])).
fof(f17583, plain, ! [X35, X34] : (~ r1(sK77(sK65(sK111)), X34) | ~ r1(X34, X35) | sP11(X35) | p103(sK65(sK111)) | ~ p102(sK65(sK111)) | ~ sP37(sK65(sK111))), inference(resolution, [], [f6156, f393])).
fof(f6156, plain, ! [X17, X15, X16] : (~ r1(sK65(sK111), X15) | ~ r1(X15, X16) | ~ r1(X16, X17) | sP11(X17)), inference(subsumption_resolution, [], [f6155, f709])).
fof(f6155, plain, ! [X17, X15, X16] : (~ r1(X15, X16) | ~ r1(sK65(sK111), X15) | ~ r1(X16, X17) | sP11(X17) | ~ sP50(sK111)), inference(subsumption_resolution, [], [f6154, f689])).
fof(f6154, plain, ! [X17, X15, X16] : (~ r1(X15, X16) | ~ r1(sK65(sK111), X15) | ~ r1(X16, X17) | sP11(X17) | ~ p101(sK111) | ~ sP50(sK111)), inference(subsumption_resolution, [], [f6106, f691])).
fof(f6106, plain, ! [X17, X15, X16] : (~ r1(X15, X16) | ~ r1(sK65(sK111), X15) | ~ r1(X16, X17) | sP11(X17) | p102(sK111) | ~ p101(sK111) | ~ sP50(sK111)), inference(resolution, [], [f5518, f319])).
fof(f5518, plain, ! [X2, X0, X3, X1] : (~ r1(sK111, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | sP11(X1)), inference(resolution, [], [f608, f695])).
fof(f608, plain, ! [X37, X35, X33, X36, X34] : (~ r1(sK110, X33) | ~ r1(X36, X37) | ~ r1(X35, X36) | ~ r1(X34, X35) | ~ r1(X33, X34) | sP11(X37)), inference(cnf_transformation, [], [f276])).
fof(f19671, plain, (~ spl121_1223 | spl121_1224 | spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_456 | ~ spl121_457 | ~ spl121_468), inference(avatar_split_clause, [], [f19654, f9796, f9723, f9719, f6580, f6468, f6463, f868, f863, f19668, f19663])).
fof(f19663, plain, (spl121_1223 <=> p104(sK90(sK89(sK77(sK65(sK111)))))), introduced(avatar_definition, [new_symbols(naming, [spl121_1223])])).
fof(f19654, plain, (p103(sK90(sK89(sK77(sK65(sK111))))) | ~ p104(sK90(sK89(sK77(sK65(sK111))))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_456 | ~ spl121_457 | ~ spl121_468)), inference(resolution, [], [f19638, f553])).
fof(f553, plain, ! [X0] : (~ sP11(X0) | p103(X0) | ~ p104(X0)), inference(cnf_transformation, [], [f226])).
fof(f19666, plain, (~ spl121_1222 | spl121_1223 | spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_456 | ~ spl121_457 | ~ spl121_468), inference(avatar_split_clause, [], [f19653, f9796, f9723, f9719, f6580, f6468, f6463, f868, f863, f19663, f19659])).
fof(f19653, plain, (p104(sK90(sK89(sK77(sK65(sK111))))) | ~ p105(sK90(sK89(sK77(sK65(sK111))))) | (spl121_17 | ~ spl121_18 | spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_456 | ~ spl121_457 | ~ spl121_468)), inference(resolution, [], [f19638, f552])).
fof(f552, plain, ! [X0] : (~ sP11(X0) | p104(X0) | ~ p105(X0)), inference(cnf_transformation, [], [f226])).
fof(f9981, plain, (~ spl121_458 | spl121_459 | ~ spl121_468), inference(avatar_split_clause, [], [f9906, f9796, f9733, f9728])).
fof(f9728, plain, (spl121_458 <=> p103(sK89(sK77(sK65(sK111))))), introduced(avatar_definition, [new_symbols(naming, [spl121_458])])).
fof(f9906, plain, (p102(sK89(sK77(sK65(sK111)))) | ~ p103(sK89(sK77(sK65(sK111)))) | ~ spl121_468), inference(resolution, [], [f9798, f488])).
fof(f488, plain, ! [X0] : (~ sP23(X0) | p102(X0) | ~ p103(X0)), inference(cnf_transformation, [], [f187])).
fof(f9980, plain, (~ spl121_457 | spl121_458 | ~ spl121_468), inference(avatar_split_clause, [], [f9905, f9796, f9728, f9723])).
fof(f9905, plain, (p103(sK89(sK77(sK65(sK111)))) | ~ p104(sK89(sK77(sK65(sK111)))) | ~ spl121_468), inference(resolution, [], [f9798, f487])).
fof(f487, plain, ! [X0] : (~ sP23(X0) | p103(X0) | ~ p104(X0)), inference(cnf_transformation, [], [f187])).
fof(f9979, plain, (spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_457), inference(avatar_contradiction_clause, [], [f9978])).
fof(f9978, plain, ($false | (spl121_148 | ~ spl121_149 | ~ spl121_167 | spl121_457)), inference(subsumption_resolution, [], [f9977, f6759])).
fof(f9977, plain, (~ sP24(sK77(sK65(sK111))) | (spl121_148 | ~ spl121_149 | spl121_457)), inference(subsumption_resolution, [], [f9976, f6470])).
fof(f9976, plain, (~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111))) | (spl121_148 | spl121_457)), inference(subsumption_resolution, [], [f9975, f6464])).
fof(f9975, plain, (p104(sK77(sK65(sK111))) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111))) | spl121_457), inference(resolution, [], [f9724, f470])).
fof(f470, plain, ! [X0] : (p104(sK89(X0)) | p104(X0) | ~ p103(X0) | ~ sP24(X0)), inference(cnf_transformation, [], [f185])).
fof(f9724, plain, (~ p104(sK89(sK77(sK65(sK111)))) | spl121_457), inference(avatar_component_clause, [], [f9723])).
fof(f9879, plain, (spl121_17 | ~ spl121_18 | ~ spl121_148), inference(avatar_contradiction_clause, [], [f9878])).
fof(f9878, plain, ($false | (spl121_17 | ~ spl121_18 | ~ spl121_148)), inference(subsumption_resolution, [], [f9877, f839])).
fof(f9877, plain, (~ sP37(sK65(sK111)) | (spl121_17 | ~ spl121_18 | ~ spl121_148)), inference(subsumption_resolution, [], [f9876, f870])).
fof(f9876, plain, (~ p102(sK65(sK111)) | ~ sP37(sK65(sK111)) | (spl121_17 | ~ spl121_148)), inference(subsumption_resolution, [], [f9875, f864])).
fof(f9875, plain, (p103(sK65(sK111)) | ~ p102(sK65(sK111)) | ~ sP37(sK65(sK111)) | ~ spl121_148), inference(resolution, [], [f6465, f395])).
fof(f395, plain, ! [X0] : (~ p104(sK77(X0)) | p103(X0) | ~ p102(X0) | ~ sP37(X0)), inference(cnf_transformation, [], [f141])).
fof(f6465, plain, (p104(sK77(sK65(sK111))) | ~ spl121_148), inference(avatar_component_clause, [], [f6463])).
fof(f9799, plain, (spl121_148 | spl121_468 | spl121_17 | ~ spl121_18 | ~ spl121_149 | ~ spl121_167), inference(avatar_split_clause, [], [f9794, f6580, f6468, f868, f863, f9796, f6463])).
fof(f9794, plain, (sP23(sK89(sK77(sK65(sK111)))) | p104(sK77(sK65(sK111))) | (spl121_17 | ~ spl121_18 | ~ spl121_149 | ~ spl121_167)), inference(subsumption_resolution, [], [f9793, f6759])).
fof(f9793, plain, (sP23(sK89(sK77(sK65(sK111)))) | p104(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111))) | (spl121_17 | ~ spl121_18 | ~ spl121_149)), inference(subsumption_resolution, [], [f9556, f6470])).
fof(f9556, plain, (sP23(sK89(sK77(sK65(sK111)))) | p104(sK77(sK65(sK111))) | ~ p103(sK77(sK65(sK111))) | ~ sP24(sK77(sK65(sK111))) | (spl121_17 | ~ spl121_18)), inference(resolution, [], [f9020, f467])).
fof(f9020, plain, (! [X17] : (~ r1(sK77(sK65(sK111)), X17) | sP23(X17)) | (spl121_17 | ~ spl121_18)), inference(subsumption_resolution, [], [f9019, f839])).
fof(f9019, plain, (! [X17] : (~ r1(sK77(sK65(sK111)), X17) | sP23(X17) | ~ sP37(sK65(sK111))) | (spl121_17 | ~ spl121_18)), inference(subsumption_resolution, [], [f9018, f870])).
fof(f9018, plain, (! [X17] : (~ r1(sK77(sK65(sK111)), X17) | sP23(X17) | ~ p102(sK65(sK111)) | ~ sP37(sK65(sK111))) | spl121_17), inference(subsumption_resolution, [], [f8982, f864])).
fof(f8982, plain, ! [X17] : (~ r1(sK77(sK65(sK111)), X17) | sP23(X17) | p103(sK65(sK111)) | ~ p102(sK65(sK111)) | ~ sP37(sK65(sK111))), inference(resolution, [], [f5903, f393])).
fof(f5903, plain, ! [X10, X11] : (~ r1(sK65(sK111), X10) | ~ r1(X10, X11) | sP23(X11)), inference(subsumption_resolution, [], [f5902, f709])).
fof(f5902, plain, ! [X10, X11] : (~ r1(sK65(sK111), X10) | ~ r1(X10, X11) | sP23(X11) | ~ sP50(sK111)), inference(subsumption_resolution, [], [f5901, f689])).
fof(f5901, plain, ! [X10, X11] : (~ r1(sK65(sK111), X10) | ~ r1(X10, X11) | sP23(X11) | ~ p101(sK111) | ~ sP50(sK111)), inference(subsumption_resolution, [], [f5853, f691])).
fof(f5853, plain, ! [X10, X11] : (~ r1(sK65(sK111), X10) | ~ r1(X10, X11) | sP23(X11) | p102(sK111) | ~ p101(sK111) | ~ sP50(sK111)), inference(resolution, [], [f5414, f319])).
fof(f5414, plain, ! [X2, X0, X1] : (~ r1(sK111, X2) | ~ r1(X2, X0) | ~ r1(X0, X1) | sP23(X1)), inference(resolution, [], [f609, f695])).
fof(f609, plain, ! [X30, X31, X29, X32] : (~ r1(sK110, X29) | ~ r1(X31, X32) | ~ r1(X30, X31) | ~ r1(X29, X30) | sP23(X32)), inference(cnf_transformation, [], [f276])).
fof(f7386, plain, (~ spl121_247 | spl121_188 | spl121_17 | ~ spl121_18 | ~ spl121_150), inference(avatar_split_clause, [], [f7381, f6473, f868, f863, f6804, f7383])).
fof(f7381, plain, (p3(sK65(sK111)) | ~ p3(sK77(sK65(sK111))) | (spl121_17 | ~ spl121_18 | ~ spl121_150)), inference(subsumption_resolution, [], [f7380, f839])).
fof(f7380, plain, (p3(sK65(sK111)) | ~ p3(sK77(sK65(sK111))) | ~ sP37(sK65(sK111)) | (spl121_17 | ~ spl121_18 | ~ spl121_150)), inference(subsumption_resolution, [], [f7379, f864])).
fof(f7379, plain, (p3(sK65(sK111)) | ~ p3(sK77(sK65(sK111))) | p103(sK65(sK111)) | ~ sP37(sK65(sK111)) | (~ spl121_18 | ~ spl121_150)), inference(subsumption_resolution, [], [f7378, f845])).
fof(f845, plain, sP44(sK65(sK111)), inference(resolution, [], [f833, f351])).
fof(f351, plain, ! [X0] : (~ sP47(X0) | sP44(X0)), inference(cnf_transformation, [], [f109])).
fof(f7378, plain, (p3(sK65(sK111)) | ~ p3(sK77(sK65(sK111))) | ~ sP44(sK65(sK111)) | p103(sK65(sK111)) | ~ sP37(sK65(sK111)) | (~ spl121_18 | ~ spl121_150)), inference(subsumption_resolution, [], [f7363, f870])).
fof(f7363, plain, (p3(sK65(sK111)) | ~ p3(sK77(sK65(sK111))) | ~ p102(sK65(sK111)) | ~ sP44(sK65(sK111)) | p103(sK65(sK111)) | ~ sP37(sK65(sK111)) | ~ spl121_150), inference(resolution, [], [f2540, f6475])).
fof(f2540, plain, ! [X17] : (~ p102(sK77(X17)) | p3(X17) | ~ p3(sK77(X17)) | ~ p102(X17) | ~ sP44(X17) | p103(X17) | ~ sP37(X17)), inference(duplicate_literal_removal, [], [f2501])).
fof(f2501, plain, ! [X17] : (p3(X17) | ~ p102(sK77(X17)) | ~ p3(sK77(X17)) | ~ p102(X17) | ~ sP44(X17) | p103(X17) | ~ p102(X17) | ~ sP37(X17)), inference(resolution, [], [f363, f393])).
fof(f363, plain, ! [X2, X0] : (~ r1(X0, X2) | p3(X0) | ~ p102(X2) | ~ p3(X2) | ~ p102(X0) | ~ sP44(X0)), inference(cnf_transformation, [], [f115])).
fof(f115, plain, ! [X0] : (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X2] : (~ p102(X2) | ~ p3(X2) | ~ r1(X0, X2)))) | ~ sP44(X0)), inference(rectify, [], [f114])).
fof(f114, plain, ! [X47] : (~ p102(X47) | ((~ p3(X47) | ! [X52] : (~ p102(X52) | p3(X52) | ~ r1(X47, X52))) & (p3(X47) | ! [X53] : (~ p102(X53) | ~ p3(X53) | ~ r1(X47, X53)))) | ~ sP44(X47)), inference(nnf_transformation, [], [f52])).
fof(f6817, plain, (~ spl121_149 | spl121_150 | ~ spl121_167), inference(avatar_split_clause, [], [f6771, f6580, f6473, f6468])).
fof(f6771, plain, (p102(sK77(sK65(sK111))) | ~ p103(sK77(sK65(sK111))) | ~ spl121_167), inference(resolution, [], [f6582, f422])).
fof(f422, plain, ! [X0] : (~ sP35(X0) | p102(X0) | ~ p103(X0)), inference(cnf_transformation, [], [f148])).
fof(f6816, plain, (spl121_17 | ~ spl121_18 | spl121_149), inference(avatar_contradiction_clause, [], [f6815])).
fof(f6815, plain, ($false | (spl121_17 | ~ spl121_18 | spl121_149)), inference(subsumption_resolution, [], [f6814, f839])).
fof(f6814, plain, (~ sP37(sK65(sK111)) | (spl121_17 | ~ spl121_18 | spl121_149)), inference(subsumption_resolution, [], [f6813, f870])).
fof(f6813, plain, (~ p102(sK65(sK111)) | ~ sP37(sK65(sK111)) | (spl121_17 | spl121_149)), inference(subsumption_resolution, [], [f6812, f864])).
fof(f6812, plain, (p103(sK65(sK111)) | ~ p102(sK65(sK111)) | ~ sP37(sK65(sK111)) | spl121_149), inference(resolution, [], [f6469, f396])).
fof(f396, plain, ! [X0] : (p103(sK77(X0)) | p103(X0) | ~ p102(X0) | ~ sP37(X0)), inference(cnf_transformation, [], [f141])).
fof(f6469, plain, (~ p103(sK77(sK65(sK111))) | spl121_149), inference(avatar_component_clause, [], [f6468])).
fof(f6733, plain, ~ spl121_17, inference(avatar_contradiction_clause, [], [f6732])).
fof(f6732, plain, ($false | ~ spl121_17), inference(subsumption_resolution, [], [f6731, f709])).
fof(f6731, plain, (~ sP50(sK111) | ~ spl121_17), inference(subsumption_resolution, [], [f6730, f689])).
fof(f6730, plain, (~ p101(sK111) | ~ sP50(sK111) | ~ spl121_17), inference(subsumption_resolution, [], [f6729, f691])).
fof(f6729, plain, (p102(sK111) | ~ p101(sK111) | ~ sP50(sK111) | ~ spl121_17), inference(resolution, [], [f865, f321])).
fof(f321, plain, ! [X0] : (~ p103(sK65(X0)) | p102(X0) | ~ p101(X0) | ~ sP50(X0)), inference(cnf_transformation, [], [f97])).
fof(f865, plain, (p103(sK65(sK111)) | ~ spl121_17), inference(avatar_component_clause, [], [f863])).
fof(f6583, plain, (spl121_17 | spl121_167 | ~ spl121_18), inference(avatar_split_clause, [], [f6578, f868, f6580, f863])).
fof(f6578, plain, (sP35(sK77(sK65(sK111))) | p103(sK65(sK111)) | ~ spl121_18), inference(subsumption_resolution, [], [f6577, f839])).
fof(f6577, plain, (sP35(sK77(sK65(sK111))) | p103(sK65(sK111)) | ~ sP37(sK65(sK111)) | ~ spl121_18), inference(subsumption_resolution, [], [f6270, f870])).
fof(f6270, plain, (sP35(sK77(sK65(sK111))) | p103(sK65(sK111)) | ~ p102(sK65(sK111)) | ~ sP37(sK65(sK111))), inference(resolution, [], [f4119, f393])).
fof(f4119, plain, ! [X5] : (~ r1(sK65(sK111), X5) | sP35(X5)), inference(subsumption_resolution, [], [f4118, f709])).
fof(f4118, plain, ! [X5] : (~ r1(sK65(sK111), X5) | sP35(X5) | ~ sP50(sK111)), inference(subsumption_resolution, [], [f4117, f689])).
fof(f4117, plain, ! [X5] : (~ r1(sK65(sK111), X5) | sP35(X5) | ~ p101(sK111) | ~ sP50(sK111)), inference(subsumption_resolution, [], [f4069, f691])).
fof(f4069, plain, ! [X5] : (~ r1(sK65(sK111), X5) | sP35(X5) | p102(sK111) | ~ p101(sK111) | ~ sP50(sK111)), inference(resolution, [], [f1453, f319])).
fof(f1453, plain, ! [X0, X1] : (~ r1(sK111, X0) | ~ r1(X0, X1) | sP35(X1)), inference(resolution, [], [f610, f695])).
fof(f610, plain, ! [X28, X26, X27] : (~ r1(sK110, X26) | ~ r1(X27, X28) | ~ r1(X26, X27) | sP35(X28)), inference(cnf_transformation, [], [f276])).
fof(f1247, plain, spl121_18, inference(avatar_contradiction_clause, [], [f1246])).
fof(f1246, plain, ($false | spl121_18), inference(subsumption_resolution, [], [f1245, f709])).
fof(f1245, plain, (~ sP50(sK111) | spl121_18), inference(subsumption_resolution, [], [f1244, f689])).
fof(f1244, plain, (~ p101(sK111) | ~ sP50(sK111) | spl121_18), inference(subsumption_resolution, [], [f1243, f691])).
fof(f1243, plain, (p102(sK111) | ~ p101(sK111) | ~ sP50(sK111) | spl121_18), inference(resolution, [], [f869, f322])).
fof(f322, plain, ! [X0] : (p102(sK65(X0)) | p102(X0) | ~ p101(X0) | ~ sP50(X0)), inference(cnf_transformation, [], [f97])).
fof(f869, plain, (~ p102(sK65(sK111)) | spl121_18), inference(avatar_component_clause, [], [f868])).