fof(f133, plain, $false, inference(subsumption_resolution, [], [f132, f40])).
fof(f40, plain, sP0(sK3), inference(resolution, [], [f31, f22])).
fof(f22, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.001.p', reflexivity)).
fof(f31, plain, ! [X1] : (~ r1(sK3, X1) | sP0(X1)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, (p100(sK3) & ~ p101(sK3) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & sP0(X1)) | ~ r1(sK3, X1)) & ! [X6] : (p2(X6) | ~ r1(sK3, X6))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3])], [f19, f20])).
fof(f20, plain, (? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & sP0(X1)) | ~ r1(X0, X1)) & ! [X6] : (p2(X6) | ~ r1(X0, X6))) => (p100(sK3) & ~ p101(sK3) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & sP0(X1)) | ~ r1(sK3, X1)) & ! [X6] : (p2(X6) | ~ r1(sK3, X6)))), introduced(choice_axiom, [])).
fof(f19, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & sP0(X1)) | ~ r1(X0, X1)) & ! [X6] : (p2(X6) | ~ r1(X0, X6))), inference(rectify, [], [f13])).
fof(f13, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & sP0(X1)) | ~ r1(X0, X1)) & ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(definition_folding, [], [f11, e12])).
fof(f12, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X6] : (p101(X6) & p2(X6) & r1(X1, X6)) & ? [X7] : (p101(X7) & ~ p2(X7) & r1(X1, X7))) | ~ sP0(X1)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X1] : (sP0(X1) <=> (~ p100(X1) | p101(X1) | (? [X6] : (p101(X6) & p2(X6) & r1(X1, X6)) & ? [X7] : (p101(X7) & ~ p2(X7) & r1(X1, X7))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f11, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p100(X1) | p101(X1) | (? [X6] : (p101(X6) & p2(X6) & r1(X1, X6)) & ? [X7] : (p101(X7) & ~ p2(X7) & r1(X1, X7))))) | ~ r1(X0, X1)) & ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(flattening, [], [f10])).
fof(f10, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & ((~ p100(X1) | p101(X1)) | (? [X6] : ((p101(X6) & p2(X6)) & r1(X1, X6)) & ? [X7] : ((p101(X7) & ~ p2(X7)) & r1(X1, X7))))) | ~ r1(X0, X1))) & ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X6] : (~ (p101(X6) & p2(X6)) | ~ r1(X1, X6)) & ~ ! [X7] : (~ (p101(X7) & ~ p2(X7)) | ~ r1(X1, X7))))) | ~ r1(X0, X1))) | ~ ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(pure_predicate_removal, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X6] : (~ (p101(X6) & ~ p102(X6) & p2(X6)) | ~ r1(X1, X6)) & ~ ! [X7] : (~ (p101(X7) & ~ p102(X7) & ~ p2(X7)) | ~ r1(X1, X7))))) | ~ r1(X0, X1))) | ~ ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(flattening, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X6] : (~ (p101(X6) & ~ p102(X6) & p2(X6)) | ~ r1(X1, X6)) & ~ ! [X7] : (~ (p101(X7) & ~ p102(X7) & ~ p2(X7)) | ~ r1(X1, X7))))) | ~ r1(X0, X1))) | ~ ! [X8] : (p2(X8) | ~ r1(X0, X8))), inference(rectify, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p2(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p2(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.001.p', main)).
fof(f132, plain, ~ sP0(sK3), inference(subsumption_resolution, [], [f131, f38])).
fof(f38, plain, p100(sK3), inference(cnf_transformation, [], [f21])).
fof(f131, plain, (~ p100(sK3) | ~ sP0(sK3)), inference(subsumption_resolution, [], [f130, f37])).
fof(f37, plain, ~ p101(sK3), inference(cnf_transformation, [], [f21])).
fof(f130, plain, (p101(sK3) | ~ p100(sK3) | ~ sP0(sK3)), inference(resolution, [], [f64, f25])).
fof(f25, plain, ! [X0] : (~ p2(sK2(X0)) | p101(X0) | ~ p100(X0) | ~ sP0(X0)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ! [X0] : (~ p100(X0) | p101(X0) | ((p101(sK1(X0)) & p2(sK1(X0)) & r1(X0, sK1(X0))) & (p101(sK2(X0)) & ~ p2(sK2(X0)) & r1(X0, sK2(X0)))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2])], [f15, f17, f16])).
fof(f16, plain, ! [X0] : (? [X1] : (p101(X1) & p2(X1) & r1(X0, X1)) => (p101(sK1(X0)) & p2(sK1(X0)) & r1(X0, sK1(X0)))), introduced(choice_axiom, [])).
fof(f17, plain, ! [X0] : (? [X2] : (p101(X2) & ~ p2(X2) & r1(X0, X2)) => (p101(sK2(X0)) & ~ p2(sK2(X0)) & r1(X0, sK2(X0)))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X0] : (~ p100(X0) | p101(X0) | (? [X1] : (p101(X1) & p2(X1) & r1(X0, X1)) & ? [X2] : (p101(X2) & ~ p2(X2) & r1(X0, X2))) | ~ sP0(X0)), inference(rectify, [], [f14])).
fof(f14, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X6] : (p101(X6) & p2(X6) & r1(X1, X6)) & ? [X7] : (p101(X7) & ~ p2(X7) & r1(X1, X7))) | ~ sP0(X1)), inference(nnf_transformation, [], [f12])).
fof(f64, plain, p2(sK2(sK3)), inference(subsumption_resolution, [], [f63, f40])).
fof(f63, plain, (~ sP0(sK3) | p2(sK2(sK3))), inference(subsumption_resolution, [], [f62, f38])).
fof(f62, plain, (~ p100(sK3) | ~ sP0(sK3) | p2(sK2(sK3))), inference(subsumption_resolution, [], [f45, f37])).
fof(f45, plain, (p101(sK3) | ~ p100(sK3) | ~ sP0(sK3) | p2(sK2(sK3))), inference(resolution, [], [f24, f30])).
fof(f30, plain, ! [X6] : (~ r1(sK3, X6) | p2(X6)), inference(cnf_transformation, [], [f21])).
fof(f24, plain, ! [X0] : (r1(X0, sK2(X0)) | p101(X0) | ~ p100(X0) | ~ sP0(X0)), inference(cnf_transformation, [], [f18])).